---
id: 68f29dbf9125476b4808d6fb
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Ahora es momento de estilizar el contenedor `settings-card`. Crea un selector de clase para `settings-card`.

Aquí es donde irá toda la estructura de formato para el contenedor. Establece el `max-width` en `250px` para definir el tamaño total de tu contenedor.

Luego, establece tu `padding` en `20px` para que tu contenido tenga espacio entre él y el borde del contenedor.

Después de esto, crea un borde redondeado estableciendo tu `border-radius` en `10px`.

Luego establece un `box-shadow` con los valores `0 2px 6px rgba(0,0,0,0.2)`. Esto creará un efecto sutil de "elevación" que dará profundidad al contenedor.

# --hints--

Debes tener un selector de clase en tu css para `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Debes agregar la propiedad `max-width` al selector de clase `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Debes tener una propiedad `max-width` con un valor de `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Debes agregar una propiedad `padding` al selector de clase `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Debes tener una propiedad `padding` con un valor de `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Debes agregar una propiedad `border-radius` a tu selector de clase `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Debes tener una propiedad `border-radius` con un valor de `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Debes agregar una propiedad `box-shadow` a tu selector de clase `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Debes tener un `box-shadow` con el valor de `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
