---
id: 579e2a2c335b9d72dd32e05c
title: Algorithmus zum Slice- und Splice-Verfahren implementieren
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

In diesem Labor müssen Sie einen Algorithmus erstellen, um zwei Arrays zusammenzuführen.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Erstellen Sie eine `frankenSplice` Funktion, die zwei Arrays und einen Index akzeptiert.
1. Kopieren Sie jedes Element des ersten Arrays in das zweite Array, in der Reihenfolge, beginnend am angegebenen Index, und geben Sie das resultierende Array zurück.
1. Die Eingabe-Arrays sollen nach Ausführung der Funktion unverändert bleiben.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` sollte `[4, 1, 2, 3, 5]` zurückgeben.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` sollte `["a", 1, 2, "b"]` zurückgeben.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` sollte `["head", "shoulders", "claw", "tentacle", "knees", "toes"]` zurückgeben.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Alle Elemente des ersten Arrays sollten in ihrer ursprünglichen Reihenfolge zum zweiten Array hinzugefügt werden. `frankenSplice([1, 2, 3, 4], [], 0)` sollte `[1, 2, 3, 4]` zurückgeben.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

Das erste Array sollte nach Ausführung der Funktion unverändert bleiben.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

Das zweite Array sollte nach Ausführung der Funktion unverändert bleiben.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
