---
id: 61329d52e5010e08d9b9d66b
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Another important `meta` element for accessibility and SEO is the `description` definition. The value of the `content` attribute is used by search engines to provide a description of your page.

Füge ein `meta`-Element hinzu, das über ein auf `description` gesetztes `name`-Attribut verfügt und weise diesem ein passendes `content`-Attribut zu.

# --hints--

Du solltest dem `head` ein neues `meta`-Element hinzufügen.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Du solltest `meta` ein, auf `description` gesetztes, `name`-Attribut zuweisen.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Du solltest `meta` ein `content`-Attribut zuweisen.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Der Wert des `content`-Attributs sollte 165 Zeichen nicht übersteigen. _Das ist die maximale Beschreibungslänge von Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
