---
id: 6153908a366afb4d57185c8d
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexbox 有主軸和交叉軸。主軸由 `flex-direction` 屬性定義，該屬性有四個可能的值：

- `row`（預設值）：從左到右排列彈性項目的水平軸
- `row-reverse`：水平軸上從右到左排列的彈性項目
- `column`：從上到下排列彈性項目的垂直軸
- `column-reverse`：垂直軸上從底部到頂部排列的彈性項目

**注意**：軸和方向會根據文字方向而不同。所顯示的數值是針對從左到右的文字方向。

嘗試不同的值以查看它們如何影響版面配置。

完成後，請在 `.gallery` 元素上設定明確的 `flex-direction` 為 `row`。

# --hints--

你的 `.gallery` 選擇器應該有一個 `flex-direction` 屬性，並將值設為 `row`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
