---
id: 61537c5f81f0cf325b4a854c
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在 `body` 元素內新增一個 `header` 元素，並為它指定一個 `header` 類別。

在 `header` 中創建一個 `h1`，文字為 `css flexbox photo gallery`。

# --hints--

你應該在你的 `body` 元素中有一個 `header` 元素。

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

你的 `header` 元素應該有一個 `class`，其值為 `header`。

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

你的 `header` 元素應該包含一個 `h1` 元素。

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

你的 `h1` 元素內應該包含文字 `css flexbox photo gallery`。

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
