---
id: 61537c9eecea6a335db6da79
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在 `.header` 元素下方，創建一個新的 `div` 元素併爲其分配 `gallery` 的 `class`。 這個 `div` 將作爲相冊的容器。

在該 `.gallery` 元素中，創建九個 `img` 元素。

# --hints--

你應該在 `body` 元素中創建一個 `div` 元素。

```js
assert.isAtLeast(document.querySelector('body')?.querySelectorAll('div')?.length, 1);
```

新創建的 `div` 元素的 `class` 應該爲 `gallery`。

```js
assert.exists(document.querySelector('body')?.querySelector('.gallery'));
```

新創建的 `div` 元素應該在 `.header` 元素後面。

```js
assert.exists(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

`.gallery` 元素內應該有九個 `img` 元素。

```js
assert.lengthOf(document.querySelector('.gallery')?.querySelectorAll('img'), 9);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
