---
id: 6153893900438b4643590367
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

將你的 `.header` 文字置中對齊。使用 `text-transform` 屬性並將值設為 `uppercase`，使文字變成大寫。

在所有邊緣給它 `32px` 的內距。將背景設為 `#0a0a23`，文字顏色設為 `#fff`。

為你的 `border-bottom` 添加值為 `4px solid #fdb347`。

# --hints--

你應該有一個 `.header` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

你的 `.header` 選擇器應該有一個 `text-align` 屬性，並將值設為 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

你的 `.header` 選擇器應該有一個 `text-transform` 屬性，並將值設為 `uppercase`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

你的 `.header` 選擇器應該有一個 `padding` 屬性，並將值設為 `32px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

你的 `.header` 選擇器應該有一個 `background-color` 屬性，並將值設為 `#0a0a23`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

你的 `.header` 選擇器應該有一個 `color` 屬性，並將值設為 `#fff`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

你的 `.header` 選擇器應該有一個 `border-bottom` 屬性，並將值設為 `4px solid #fdb347`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
