---
id: 6564d97a5ef99783216a6229
title: Hatua 44
challengeType: 20
dashedName: step-44
---

# --description--

Python hutoa aina maalum ya mfuatano wa herufi unaoitwa mfuatano wa herufi *mbichi* (raw string). Mfuatano wa herufi mbichi huanza na kiambishi awali cha `r`. Tofauti kuu na mfuatano wa herufi wa kawaida iko katika jinsi wanavyoshughulikia herufi ya mkwaju wa nyuma: katika mfuatano wa herufi mbichi, mikwaju wa nyuma hutambuliwa kama herufi halisi badala ya kuwa herufi za kutoroka. Unapoandika usemi wa kawaida, kutumia mfuatano wa herufi mbichi ni tabia nzuri, kwa sababu mara nyingi hutumia herufi nyingi za `\`.

Badilisha mfuatano wako wa herufi wa `pattern` kuwa mfuatano wa herufi mbichi kwa kuanza na kiambishi awali cha `r`.

# --hints--

Unapaswa kubadilisha kigezo chako cha `pattern` kuwa `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
