---
id: 65648f4b2281ba50051ae39c
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Ifuatayo, andika mzunguko wa `for` ukiwa na `i` kama kigezo cha mzunguko. Tumia kitendakazi cha `range()` kurudia hadi thamani ya `length`.

Ndani ya mzunguko, tumia opereta ya kuongeza kuongeza herufi ya bahati nasibu kutoka `all_characters` kwa thamani ya sasa ya `password`. Tumia kitendakazi cha `choice()` kutoka moduli ya `secrets` kwa ajili ya hilo.

# --hints--

Unapaswa kuandika mzunguko wa `for` unaorudia juu ya `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Unapaswa kutumia opereta ya `+=` kuongeza herufi ya bahati nasibu kutoka `all_characters` kwa thamani ya sasa ya `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
