---
id: 65646ab6bced52360f44670f
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Sasa tangaza kigezo kinachoitwa `all_characters` na weka thamani yake kuwa matokeo ya kuunganisha vigezo vyako vilivyopo.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `all_characters`.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Unapaswa kuunganisha `letters`, `digits`, na `symbols` na kuweka matokeo hayo katika kigezo chako cha `all_characters`.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
