---
id: 656470d517833a39bb8b5608
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Moduli ya `random` ina kizalishaji cha nambari za nasibu za bandia. Kazi nyingi zake hutegemea kitendakazi cha `random()`, ambacho hurudisha nambari ya desimali katika anuwai kati ya `0.0` (ikiwa ni pamoja na) na `1.0` (isipokuwa).

Fanya wito wa kitendakazi cha `random()` kutoka moduli ya `random` na chapisha matokeo.

# --hints--

Unapaswa kuchapisha matokeo ya wito la `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
