---
id: 656469e62377ab34a5bcb8d1
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Tangaza vigezo viwili vipya vinavyoitwa `digits` na `symbols` na weka thamani `string.digits` na `string.punctuation` kwao, kwa mpangilio huo.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Unapaswa kuweka thamani `string.digits` kwa kigezo chako cha `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Unapaswa kutangaza kigezo kinachoitwa `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Unapaswa kuweka thamani `string.punctuation` kwa kigezo chako cha `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
