---
id: 656475bbf1c2573de1d2c69c
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Kila wakati msimbo unapoendesha, unapaswa kuona herufi isiyotarajiwa kutoka kwenye mfuatano wa herufi wa `all_characters`. Hii ndilo hasa unalotaka kufanikisha kuunda nenosiri lisilotabirika.

Hata hivyo, algoriti ambayo `random` inategemea hufanya nambari bandia zisizotarajiwa kuwa za kutabirika. Kwa hiyo, ingawa moduli ya `random` ni inayofaa kwa programu nyingi za kawaida, haiwezi kutumika kwa madhumuni ya usimbaji fiche, kutokana na asili yake ya kutabirika.

Badala ya kuimport `random`, import moduli ya `secrets`. Kisha badilisha wito wa `print()` ili utumie `secrets.choice(all_characters)`.

# --hints--

Unapaswa kuimport moduli ya `secrets` badala ya moduli ya `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Unapaswa kubadilisha wito wako wa `print()` uliopo ili uchapishe `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
