---
id: 65646ffeaed2d238c562a014
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Ni mkataba wa kawaida kuweka tamko za `import` juu ya msimbo wako. Zaidi ya hayo, ikiwa kuna tamko nyingi za `import`, zipangilie kwa mpangilio wa alfabeti ili kuboresha usomaji.

Juu ya msimbo wako, import moduli ya `random`.

# --hints--

Unapaswa kuimport moduli ya `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
