---
id: 656485a8a3496d4a36b1496a
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Hizi ni tofauti tatu zinazounda herufi zinazowezekana kuchaguliwa wakati wa kuunda nenosiri.

Mbele yao kidogo, ongeza maoni yanayosema `Define the possible characters for the password`.

# --hints--

Unapaswa kuongeza maoni hayo juu kidogo ya tofauti ya `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
