---
id: 6564cbb6311a40783b3f5de6
title: Hatua 36
challengeType: 20
dashedName: step-36
---

# --description--

Darasa la vitu la herufi pia linakuwezesha kuonyesha safu ya herufi zinazolingana. Unahitaji kubainisha herufi ya kuanzia na herufi ya mwisho zilizoachwa na alama ya msalaba, `-`. Herufi hufuata mpangilio wa Unicode.

Badilisha tofauti yako ya `pattern` ili iendane na herufi yoyote `t` inayotanguliwa na herufi ndogo katika tofauti ya `quote`. Tumia safu ya herufi kutoka `a` hadi `z` kwa ajili hiyo.

# --hints--

Unapaswa kubadilisha tofauti yako ya `pattern` ili iendane na herufi yoyote `t` katika `quote` inayotanguliwa na herufi ndogo kwa kutumia darasa la vitu la `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
