---
id: 65648e4c5b316c4ec5e4fddc
title: Hatua 15
challengeType: 20
dashedName: step-15
---

# --description--

Chini kabisa ya kitendakazi chako, tangaza mabadiliko ya `password` na weka thamani ya mfuatano wa herufi tupu kwa mabadiliko haya.

# --hints--

Unapaswa kuwa na mabadiliko ya `password` ndani ya kitendakazi chako yenye thamani ya mfuatano wa herufi tupu. Zingatia ulinganifu wa mistari.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
