---
id: 6564642ce403fe29cc23fee7
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Moduli ya Python ni faili yenye msimbo ulioundwa kutekeleza kazi maalum. Maktaba ya kawaida ya Python ina moduli nyingi ambazo unaweza ku-import na kutumia katika msimbo wako. Unaweza kufanikisha hili kwa kutumia kauli ya `import` ikifuatiwa na jina la moduli.

Anza mradi huu kwa ku-import moduli ya `string`.

# --hints--

Unapaswa ku-import moduli ya `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
