---
id: 656474314bf37d3c83a85143
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Kitendakazi cha `choice()` kutoka moduli ya `random` kinachukua mfululizo na kurudisha kipengee chochote cha mfululizo hicho kwa bahati nasibu.

Badilisha wito wako wa `print()` kutumia kitendakazi cha `choice()` na upitie `all_characters` kama hoja.

# --hints--

Unapaswa kubadilisha wito wako uliopo wa `print()` ili uchapishe `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
