---
id: 65649a631440b757f9f7b428
title: Hatua 28
challengeType: 20
dashedName: step-28
---

# --description--

Usemi wa kawaida, au regex, ni muundo unaotumika kulinganisha mchanganyiko maalum wa herufi ndani ya mfuatano wa herufi. Ni mbadala halali kwa taarifa za masharti za `if`/`else` unapo hitaji kulinganisha au kutafuta mifumo ndani ya mfuatano wa herufi kwa madhumuni ya uthibitisho, kubadilisha herufi, na mengineyo.

Kitendakazi cha `compile()` kutoka moduli ya `re` hukusanya mfuatano wa herufi uliotumwa kama hoja kuwa kitu cha usemi wa kawaida ambacho kinaweza kutumika na njia nyingine za `re`.

Tangaza kigezo kipya cha `pattern` na weka thamani ya `re.compile('i')` kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo cha `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Unapaswa kuweka `re.compile('i')` kwa kigezo chako cha `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
