---
id: 6564f834dd717998092cfd47
title: Hatua 54
challengeType: 20
dashedName: step-54
---

# --description--

Baada ya maoni yako mapya, andika mzunguko wa `for` rudia juu ya orodha ya `constraints`. Tumia `constraint` na `pattern` kama vigezo vya mzunguko.

# --hints--

Unapaswa kuandika mzunguko wa `for` rudia juu ya orodha ya `constraints`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("generate_password").find_whiles()[0].find_bodies()[0].find_for_loops()[1].find_for_iter().is_equivalent("constraints")`)) })
```

Mzunguko wako wa `for` unapaswa kutumia `constraint` na `pattern` kama vigezo vya mzunguko rudia juu ya orodha ya `constraints`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("generate_password").find_whiles()[0].find_bodies()[0].find_for_loops()[1].find_for_vars().is_equivalent("constraint, pattern")`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
--fcc-editable-region--         
        # Check constraints
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
