---
id: 6564683821b2ee3174e7250e
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Unaweza kupata huduma zilizotangazwa ndani ya moduli iliyoleweshwa kupitia uandishi wa nukta. Uandishi wa nukta hufanya kazi kwa kuongeza nukta ikifuatiwa na jina la huduma kwa jina la moduli. Kwa mfano, hivi ndivyo unavyopata thabiti ya `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

Katika mradi huu, utatumia thabiti tofauti kutoka moduli ya `string`.

Tangaza kigezo kipya kinachoitwa `letters` na weka thamani ya `string.ascii_letters` kwa kigezo hiki.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Unapaswa kuweka thamani ya `string.ascii_letters` kwa kigezo chako cha `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
