---
id: 6564c67db26c417561ab510d
title: Hatua 31
challengeType: 20
dashedName: step-31
---

# --description--

Kama unavyoona kutoka kwenye matokeo, sasa usemi wako wa kawaida unalingana na `l` ya kwanza ndani ya mfuatano wa herufi.

Katika muundo wako, unaweza kuongeza kihimarisha baada ya herufi ili kubainisha ni mara ngapi herufi hiyo inapaswa kurudiwa. Kwa mfano, kihimarisha cha `+` kinamaanisha inapaswa kurudiwa mara moja au zaidi.

Ongeza kihimarisha cha `+` kwenye muundo wako.

# --hints--

Unapaswa kubadilisha variable yako ya `pattern` kuwa `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
