---
id: 656477845006313fbfea0ad1
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Ingawa athari inaweza kuonekana sawa na `random.choice()`, `secrets` inakuhakikishia chanzo salama zaidi cha nasibu ambacho mfumo wako wa uendeshaji unaweza kutoa.

Sasa, futa wito wako wa `print()` wawili.

# --hints--

Unapaswa kufuta wito wako wa `print()` wawili.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
