---
id: 6564883669b5af4b69f794cc
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Kigezo chako cha `all_characters` ni mfuatano wa herufi ulioundwa na herufi ndogo na herufi kubwa zote, nambari zote 10 na herufi maalum kadhaa.

Mbele yake, ongeza maoni yanayosema `Combine all characters`.

# --hints--

Unapaswa kuongeza maoni hayo juu kidogo ya kigezo chako cha `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
