---
id: 6564b8c9349bd76dc037967b
title: Hatua 29
challengeType: 20
dashedName: step-29
---

# --description--

Kitendakazi cha `search()` kutoka moduli ya `re` huchambua mfuatano wa herufi uliotumwa kama hoja kwa kutafuta mahali pa kwanza ambapo muundo wa usemi wa kawaida unalingana na mfuatano huo.

Tangaza kigezo kinachoitwa `quote` na weka thamani ya mfuatano wa herufi `'Not all those who wander are lost.'` katika kigezo hiki. Kisha, chapisha matokeo ya `pattern.search(quote)`.

# --hints--

Unapaswa kuwa na kigezo cha `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Unapaswa kuweka thamani ya mfuatano wa herufi uliotolewa katika kigezo chako kipya cha `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Unapaswa kuchapisha `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
