---
id: 656ef5bc5c0cd464be1df675
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Ifuatayo, utatoa vigezo zaidi kwa kitendakazi chako ambavyo vitatumika kama vizingiti kwa nenosiri linalotengenezwa.

Badilisha tamko la kitendakazi chako kwa kuongeza `nums`, `special_chars`, `uppercase`, na `lowercase` kwa mpangilio huu baada ya kigezo kilichopo cha `length`.

# --hints--

Kitendakazi chako kinapaswa kuchukua `length`, `nums`, `special_chars`, `uppercase`, na `lowercase` kama vigezo. Mpangilio ni muhimu.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
