---
id: 656eec2f40d18056cc58b229
title: Hatua 30
challengeType: 20
dashedName: step-30
---

# --description--

Thamani ya `None` inarudishwa kwa sababu `'i'` haipatikani ndani ya mfuatano wa herufi uliotafsiriwa.

Sasa, badilisha mfuatano wa herufi unaopitishwa kwa `re.compile()` kuwa `'l'` na uone matokeo.

# --hints--

Unapaswa kubadilisha kigezo chako cha `pattern` kuwa `re.compile('l')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)    
--fcc-editable-region--
pattern = re.compile('i')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
