---
id: 656f0d6bc60b58b9777a36c9
title: Hatua 34
challengeType: 20
dashedName: step-34
---

# --description--

Darasa la vitu linaonyeshwa kwa mabano ya mraba na linaendana na tabia moja kati ya zile zilizotajwa kati ya mabano hayo. Kwa mfano, `ma[dnt]` linaendana na `mad`, `man`, au `mat`.

Badilisha muundo wako ili uendane na `w` ikifuatiwa na `h` au `a`.

# --hints--

Muundo wako unapaswa kutumia darasa la vitu kuendana na `w` ikifuatiwa na `h` au `a`.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
