---
id: 657dadf4d8b93c1704f3a57c
title: Hatua 66
challengeType: 20
dashedName: step-66
---

# --description--

Inafanya kazi, lakini bado kuna mambo machache unaweza kuboresha. Kwanza kabisa, kuita kitendakazi kwa hoja 5 kunaweza kuleta mkanganyiko kuhusu ni thamani gani itapewa kigezo gani.

Unaweza kuita kitendakazi kwa kutumia hoja za maneno, yaani kuandika jina la kigezo wazi kisha kufuatwa na opereta ya mgawo na thamani. Kwa mfano:

```py
def add(x, y):
    return x + y

add(x=1, y=7) # 8
```

Badilisha wito lako la kitendakazi ili utumie hoja za maneno.

# --hints--

Kigezo chako cha `new_password` kinapaswa kuwa na thamani ya `generate_password(length=8, nums=1, special_chars=1, uppercase=1, lowercase=1)`.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*length\s*=\s*8\s*,\s*nums\s*=\s*1\s*,\s*special_chars\s*=\s*1\s*,\s*uppercase\s*=\s*1\s*,\s*lowercase\s*=\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
   
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (special_chars, fr'[{symbols}]'),
            (uppercase, r'[A-Z]'),
            (lowercase, r'[a-z]')
        ]

        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break
    
    return password
    
--fcc-editable-region-- 
new_password = generate_password(8, 1, 1, 1, 1)
print(new_password)
--fcc-editable-region--
```
