---
id: 66bba6fff611169359d9d36a
title: 建置機密電子郵件頁面
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `main` 元素，其 `id` 為 `email`。
1. 你的 `#email` 元素應該有 `50px` 的 `padding`、`50px` 的上邊距、`500px` 的 `width`，以及 `2px` 的 `border`。
1. 你的 `#email` 元素的總寬度，包括內距和邊框，應該是 `500px`。
1. 你應該在你的 `#email` 元素中有兩個 `div` 元素，一個的 id 為 `confidential`，另一個的 id 為 `top-secret`。
1. 你的 `#confidential` 和 `#top-secret` 元素應該有 `inline-block` 的 `display`。
1. 你的 `#confidential` 和 `#top-secret` 元素應該有 `padding`、左邊距和 `border`。
1. `#confidential` 元素應該有文字 `CONFIDENTIAL`。
1. `#top-secret` 元素應該有文字 `TOP SECRET`。
1. 你的 `#confidential` 和 `#top-secret` 元素應該使用 CSS transform 進行旋轉。
1. 你的 `#email` 元素中應該至少有三個段落元素。
1. 你的段落元素中應該至少有三個帶有 `blurred` 類別的 `span` 元素。
1. 你應該有一個 `blurred` 選擇器，使用 CSS filter 對元素進行 `3px` 模糊處理。

# --hints--

你應該有一個 `main` 元素，其 `id` 為 `email`。

```js
assert.exists(document.querySelector('main#email'));
```

你應該有一個 `#email` 選擇器，將其元素的所有邊的 `padding` 設為 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

你應該有一個 `#email` 選擇器，將其元素的 `margin-top` 設為 `50px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

你應該有一個 `#email` 選擇器，將其元素的 `width` 設定為 500px。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

你的 `#email` 元素應該有一個 `2px` 邊框。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

你的 `#email` 元素的 `box-sizing` 應該是 `border-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

你應該在你的 `#email` 元素中有一個 `id` 為 `confidential` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

你應該在你的 `#email` 元素中有一個 `id` 為 `top-secret` 的 `div` 元素。

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

你的 `#confidential` 元素的 `display` 應設定為 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

你的 `#top-secret` 元素應該將其 `display` 設定為 `inline-block`。

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

你的 `#confidential` 元素應該在所有邊緣都有 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#top-secret` 元素應該在所有邊緣都有 `padding`。

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

你的 `#confidential` 元素應該有左邊距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

你的 `#top-secret` 元素應該有左邊距。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

你的 `#confidential` 元素應該有一個 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

你的 `#top-secret` 元素應該有一個 `border`。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

你的 `#confidential` 元素的文字應該是 `CONFIDENTIAL`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

你的 `#top-secret` 元素的文字應該是 `TOP SECRET`。你要麼遺漏了文字，要麼有拼寫錯誤。

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

你的 `#confidential` 元素應該使用 CSS `transform` 來旋轉該元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

你的 `#top-secret` 元素應該使用 CSS `transform` 來旋轉該元素。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

你的 `#email` 元素中應該至少有三個段落元素。

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

你的段落中應該至少有三個帶有 `blurred` 類別的 `span` 元素。

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

你的 `.blurred` 元素不應該是空的。

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

你應該有一個 `.blurred` 選擇器，將其元素的 `filter` 設為 `blur(3px)`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
