---
id: 587d8248367417b2b2512c3c
title: helmet.hsts()로 브라우저가 HTTPS로만 사이트에 접근하도록 요청하기
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 구축되고 있습니다.

HTTP Strict Transport Security(HSTS)는 프로토콜 다운그레이드 공격과 쿠키 탈취로부터 웹사이트를 보호하는 웹 보안 정책입니다. 웹사이트에 HTTPS로 접근할 수 있다면 사용자의 브라우저가 안전하지 않은 HTTP 사용을 피하도록 요청할 수 있습니다. Strict-Transport-Security 헤더를 설정하면, 지정한 기간 동안 브라우저가 이후 요청에 대해 HTTPS를 사용하도록 지시합니다. 이는 초기 요청 이후에 오는 요청에 적용됩니다.

# --instructions--

`helmet.hsts()`를 다음 90일 동안 HTTPS를 사용하도록 설정하세요. 설정 객체 `{maxAge: timeInSeconds, force: true}`를 전달하세요. `ninetyDaysInSeconds = 90*24*60*60;`에 사용할 변수 `timeInSeconds`를 생성할 수 있습니다.

참고: 맞춤형 웹사이트에서 HTTPS를 구성하려면 도메인과 SSL/TLS 인증서가 필요합니다.

# --hints--

helmet.hsts() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge는 7776000초(90일)와 같아야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```
