---
id: 587d8247367417b2b2512c36
title: Helmet 설치 및 require하기
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

이 도전 과제들은 다음 방법 중 하나를 사용해 코드를 작성하는 작업을 포함합니다:

- <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">이 GitHub 저장소</a>를 클론하여 로컬에서 도전 과제를 완료하세요.
- 원하는 사이트 빌더를 사용해 프로젝트를 완성하세요. GitHub 저장소의 모든 파일을 반드시 포함하세요.

Helmet은 다양한 HTTP 헤더를 설정해 Express 앱을 보호하는 데 도움을 줍니다.

# --instructions--

이 수업의 모든 코드는 `myApp.js` 파일에서, 저희가 시작해둔 코드 라인 사이에 작성하세요. 저희가 추가한 코드는 변경하거나 삭제하지 마세요.

Helmet 버전 `3.21.3`이 이미 설치되어 있으니, `helmet`에서 `myApp.js`로 require하세요.

# --hints--

`helmet` 버전 `3.21.3`가 `package.json`에 있어야 합니다.

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
