---
id: 587d8247367417b2b2512c37
title: helmet.hidePoweredBy()로 잠재적으로 위험한 정보 숨기기
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있습니다.

해커는 사이트가 Express로 구동된다는 사실을 알게 되면 Express/노드의 알려진 취약점을 악용할 수 있습니다. `X-Powered-By: Express`는 기본적으로 Express에서 오는 모든 요청에 포함되어 전송됩니다. `helmet.hidePoweredBy()` 미들웨어를 사용해 X-Powered-By 헤더를 제거하세요.

# --hints--

helmet.hidePoweredBy() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```
