---
id: 587d8249367417b2b2512c3e
title: helmet.noCache()로 클라이언트 측 캐싱 비활성화
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있습니다.

웹사이트 업데이트를 배포할 때 사용자가 항상 최신 버전을 다운로드하도록 하려면 클라이언트 브라우저에서 캐싱을 비활성화할 수 있습니다(시도할 수 있습니다). 개발 중에도 유용할 수 있습니다. 캐싱은 성능 향상에 도움이 되므로 실제로 필요할 때만 이 옵션을 사용하세요.

# --instructions--

서버에서 `helmet.noCache()` 메서드를 사용하세요.

# --hints--

helmet.noCache() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
