---
id: 587d8248367417b2b2512c3a
title: helmet.noSniff()로 응답 MIME 유형 추론 방지
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있습니다.

브라우저는 콘텐츠 또는 MIME 스니핑을 사용해 응답의 `Content-Type` 헤더를 무시하고 암묵적인 콘텐츠 유형으로 데이터를 추측하여 처리할 수 있습니다. 이런 방식이 일부 상황에서는 편리할 수 있지만, 위험한 공격으로 이어질 수도 있습니다. 이 미들웨어는 `X-Content-Type-Options` 헤더를 `nosniff`로 설정하여 브라우저가 제공된 `Content-Type`를 우회하지 않도록 지시합니다.

# --instructions--

서버에서 `helmet.noSniff()` 메서드를 사용하세요.

# --hints--

helmet.noSniff() 미들웨어가 올바르게 장착되어야 합니다.

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```
