---
id: 587d8248367417b2b2512c3d
title: helmet.dnsPrefetchControl()로 DNS 사전 가져오기 비활성화
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있음을 상기하세요.

성능 향상을 위해 대부분 브라우저는 페이지 내 링크의 DNS 레코드를 사전 가져오기(prefetch)합니다. 이렇게 하면 사용자가 링크를 클릭할 때 대상 IP가 이미 알려져 있습니다. 이로 인해 DNS 서비스 과다 사용(수백만 명이 방문하는 대형 웹사이트 소유 시), 개인정보 문제(도청자가 특정 페이지에 접속 중임을 추론할 수 있음), 페이지 통계 왜곡(방문하지 않은 링크가 방문한 것으로 나타날 수 있음) 등이 발생할 수 있습니다. 보안이 매우 중요한 경우 성능 저하를 감수하고 DNS 사전 가져오기를 비활성화할 수 있습니다.

# --instructions--

서버에서 `helmet.dnsPrefetchControl()` 메서드를 사용하세요.

# --hints--

helmet.dnsPrefetchControl() 미들웨어가 올바르게 장착되어야 합니다.

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
