---
id: 587d8247367417b2b2512c38
title: helmet.frameguard()로 클릭재킹 위험 완화하기
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있습니다.

사용자의 동의 없이 페이지가 `<frame>` 또는 `<iframe>`에 삽입될 수 있습니다. 이로 인해 클릭재킹 공격 등이 발생할 수 있습니다. 클릭재킹은 사용자가 생각하는 페이지와 다른 페이지와 상호작용하도록 속이는 기법입니다. 이는 iframe을 통해 악의적인 환경에서 페이지를 실행함으로써 이루어질 수 있습니다. 그런 환경에서 해커는 페이지 위에 숨겨진 레이어를 올릴 수 있습니다. 숨겨진 버튼을 사용해 악성 스크립트를 실행할 수 있습니다. 이 미들웨어는 X-Frame-Options 헤더를 설정합니다. 누가 사이트를 프레임에 넣을 수 있는지 제한합니다. 세 가지 모드가 있습니다: DENY, SAMEORIGIN, ALLOW-FROM.

우리 앱은 프레임에 삽입될 필요가 없습니다.

# --instructions--

구성 객체 `helmet.frameguard()`를 전달하여 `{action: 'deny'}`를 사용하세요.

# --hints--

helmet.frameguard() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

helmet.frameguard()의 'action'은 'DENY'로 설정되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```
