---
id: 587d8249367417b2b2512c3f
title: helmet.contentSecurityPolicy()로 콘텐츠 보안 정책 설정하기
challengeType: 2
forumTopicId: 301585
dashedName: set-a-content-security-policy-with-helmet-contentsecuritypolicy
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

이 도전 과제는 현대 브라우저에서 여러 유형의 공격 위험과 영향을 크게 줄일 수 있는 유망한 새로운 방어책을 강조합니다. 콘텐츠 보안 정책을 설정하고 구성하면 페이지에 의도하지 않은 내용이 주입되는 것을 방지할 수 있습니다. 이를 통해 앱을 XSS 취약점, 원치 않는 추적, 악성 프레임 등으로부터 보호할 수 있습니다. CSP는 신뢰할 수 있는 콘텐츠 출처의 허용 목록을 정의하는 방식으로 작동합니다. 웹 페이지가 필요로 하는 각 리소스 종류(스크립트, 스타일시트, 폰트, 프레임, 미디어 등)에 대해 이를 구성할 수 있습니다. 여러 지시문이 있어 웹사이트 소유자가 세밀하게 제어할 수 있습니다. 자세한 내용은 HTML 5 Rocks, KeyCDN을 참조하세요. 안타깝게도 CSP는 구형 브라우저에서는 지원되지 않습니다.

기본적으로 지시문은 매우 개방적이므로, 대체 수단으로 defaultSrc 지시문을 설정하는 것이 중요합니다. Helmet은 defaultSrc와 default-src 두 가지 명명 방식을 모두 지원합니다. 대체 수단은 대부분 명시되지 않은 지시문에 적용됩니다.

# --instructions--

이번 연습에서는 `helmet.contentSecurityPolicy()`를 사용하세요. `directives` 객체를 추가하여 구성하세요. 객체 내에서 `defaultSrc`를 `["'self'"]`로 설정하세요(허용된 출처 목록은 배열이어야 합니다). 이렇게 하면 기본적으로 웹사이트 주소만 신뢰하게 됩니다. 또한 `scriptSrc` 지시문을 설정하여 스크립트가 웹사이트(`'self'`)와 도메인 `'trusted-cdn.com'`에서만 다운로드되도록 허용하세요.

힌트: `'self'` 키워드에서 작은따옴표는 키워드 자체의 일부이므로, 작동하려면 큰따옴표로 감싸야 합니다.

# --hints--

helmet.contentSecurityPolicy() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'csp');
```

csp 구성이 올바르지 않습니다. defaultSrc는 ["'self'"]이어야 하고 scriptSrc는 ["'self'", 'trusted-cdn.com']이어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
var cspHeader = Object.keys(data.headers).filter(function (k) {
  return (
    k === 'content-security-policy' ||
    k === 'x-webkit-csp' ||
    k === 'x-content-security-policy'
  );
})[0];
assert.equal(
  data.headers[cspHeader],
  "default-src 'self'; script-src 'self' trusted-cdn.com"
);
```
