---
id: 587d8249367417b2b2512c40
title: ‘parent’ helmet() 미들웨어를 사용해 Helmet 구성하기
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 구축되고 있습니다.

`app.use(helmet())`는 위에서 소개한 모든 미들웨어를 자동으로 포함하지만 `noCache()`과 `contentSecurityPolicy()`는 제외하며, 필요에 따라 이들을 활성화할 수 있습니다. 또한 구성 객체를 사용해 다른 미들웨어를 개별적으로 비활성화하거나 구성할 수도 있습니다.

**예시:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

학습 목적과 테스트 편의를 위해 각 미들웨어를 별도로 소개했습니다. ‘parent’ `helmet()` 미들웨어를 사용하는 것은 실제 프로젝트에서 쉽게 구현할 수 있습니다.

# --hints--

테스트 없음 - 설명용 도전 과제입니다

```js
assert(true);
```
