---
id: 587d8247367417b2b2512c39
title: helmet.xssFilter()로 교차 사이트 스크립팅(XSS) 공격 위험 완화하기
challengeType: 2
forumTopicId: 301583
dashedName: mitigate-the-risk-of-cross-site-scripting-xss-attacks-with-helmet-xssfilter
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 구축되고 있습니다.

교차 사이트 스크립팅(XSS)은 취약한 페이지에 악성 스크립트를 주입하여 세션 쿠키나 비밀번호 같은 민감한 데이터를 훔치는 것을 목적으로 하는 빈번한 공격 유형입니다.

XSS 공격 위험을 낮추는 기본 규칙은 간단합니다: "사용자의 입력을 절대 신뢰하지 마세요". 개발자는 외부에서 들어오는 모든 입력을 항상 정화해야 합니다. 여기에는 폼에서 오는 데이터, GET 쿼리 URL, 심지어 POST 본문에서 오는 데이터도 포함됩니다. 정화란 위험할 수 있는 문자들, 예를 들어 &lt;, > 같은 문자를 찾아 인코딩하는 것을 의미합니다.

최신 브라우저는 더 나은 소프트웨어 전략을 채택하여 위험을 완화하는 데 도움을 줄 수 있습니다. 이러한 전략은 종종 HTTP 헤더를 통해 구성할 수 있습니다.

X-XSS-Protection HTTP 헤더는 기본적인 보호 기능입니다. 브라우저가 휴리스틱 필터를 사용해 잠재적인 주입 스크립트를 감지합니다. 헤더가 활성화되어 있으면 브라우저가 스크립트 코드를 변경하여 무력화합니다. 다만 지원 범위는 제한적입니다.

# --instructions--

서버로 전송되는 입력을 정화하려면 `helmet.xssFilter()`를 사용하세요.

# --hints--

helmet.xssFilter() 미들웨어가 올바르게 장착되어 있어야 합니다.

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'xXssProtection');
assert.property(data.headers, 'x-xss-protection');
```
