---
id: 587d8248367417b2b2512c3b
title: helmet.ieNoOpen()으로 IE가 신뢰하지 않는 HTML을 열지 못하게 하기
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

참고로, 이 프로젝는 <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 스타터 프로젝트를 기반으로 만들어지고 있습니다.

일부 웹 애플리케이션은 신뢰하지 않는 HTML을 다운로드용으로 제공합니다. Internet Explorer의 일부 버전은 기본적으로 해당 HTML 파일을 사이트의 컨텍스트에서 엽니다. 이는 신뢰하지 않는 HTML 페이지가 사용자의 페이지 컨텍스트에서 악의적인 행동을 시작할 수 있다는 뜻입니다. 이 미들웨어는 X-Download-Options 헤더를 noopen으로 설정합니다. 이렇게 하면 IE 사용자가 신뢰하는 사이트 컨텍스트에서 다운로드를 실행하지 못하게 됩니다.

# --instructions--

서버에서 `helmet.ieNoOpen()` 메서드를 사용하세요.

# --hints--

helmet.ieNoOpen() 미들웨어가 올바르게 장착되어야 합니다

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
