---
id: 58a25bcef9fc0f352b528e7c
title: BCrypt 해시 이해하기
challengeType: 2
forumTopicId: 301586
dashedName: understand-bcrypt-hashes
---

# --description--

다음 도전 과제에서는 이전과 다른 새로운 시작 프로젝트를 사용하게 됩니다. 복제할 새로운 시작 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 찾을 수 있습니다.

BCrypt 해시는 매우 안전합니다. 해시는 기본적으로 원본 데이터의 지문과 같아서 항상 고유합니다. 이는 원본 데이터를 알고리즘에 입력해 고정 길이 결과를 반환하는 방식으로 이루어집니다. 이 과정을 더 복잡하고 안전하게 만들기 위해 해시에 *솔트*를 추가할 수 있습니다. 해시에 솔트를 추가한다는 것은 해싱 과정 전에 원본 데이터에 무작위 데이터를 더하는 것으로, 해시를 깨기 훨씬 어렵게 만듭니다.

BCrypt 해시는 항상 `$2a$13$ZyprE5MRw2Q3WpNOGZWGbeG7ADUre1Q8QO.uUUtcbqloU0yvzavOm`처럼 보이며 구조가 있습니다. 처음의 작은 데이터 `$2a`는 어떤 해시 알고리즘이 사용되었는지 정의합니다. 다음 부분 `$13`는 *비용(cost)*을 정의합니다. 비용은 해시를 계산하는 데 필요한 연산량을 의미합니다. 2^cost의 로그 스케일로, 데이터가 해싱 알고리즘을 몇 번 거치는지 결정합니다. 예를 들어, 비용이 10일 때는 평균 컴퓨터에서 초당 10개의 비밀번호를 해시할 수 있지만, 비용이 15일 때는 해시 하나당 3초가 걸립니다... 더 나아가 비용이 31이면 해시를 완료하는 데 며칠이 걸릴 것입니다. 현재 비용 12는 매우 안전한 수준으로 간주됩니다. 해시의 마지막 부분 `$ZyprE5MRw2Q3WpNOGZWGbeG7ADUre1Q8QO.uUUtcbqloU0yvzavOm`는 숫자, 점, 문자로 이루어진 긴 문자열처럼 보이지만 실제로는 두 개의 별도 정보입니다. 처음 22자는 평문 솔트이고, 나머지는 해시된 비밀번호입니다!

# --instructions--

이 수업들의 모든 코드는 `server.js` 파일에서 우리가 시작해 둔 코드 사이에 작성하세요. 우리가 추가한 코드는 변경하거나 삭제하지 마세요.

BCrypt는 이미 의존성으로 추가되어 있으니, 서버에서 `bcrypt`로 require 하세요.

정답이라고 생각되면 페이지를 제출하세요.

# --hints--

BCrypt는 의존성으로 있어야 합니다.

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
var packJson = JSON.parse(data);
assert.property(
  packJson.dependencies,
  'bcrypt',
  'Your project should list "bcrypt" as a dependency'
);
```

BCrypt가 제대로 require 되어야 합니다.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /bcrypt.*=.*require.*('|")bcrypt('|")/gi,
  'You should correctly require and instantiate bcrypt as bcrypt.'
);
```
