---
id: 58a25bcff9fc0f352b528e7d
title: 비동기로 비밀번호 해시 및 비교하기
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 시작 프로젝트를 기반으로 하고 있습니다.

해시는 계산 집약적으로 설계되었기 때문에, 해시하는 동안 들어오는 연결이 차단되지 않도록 서버에서 비동기로 처리하는 것이 권장됩니다. 비밀번호를 비동기로 해시하려면 단순히 다음을 호출하면 됩니다.

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

이 해시 함수를 서버에 추가하세요(함수에서 사용할 변수를 이미 정의해 두었습니다) 그리고 콘솔에 출력하여 확인하세요! 이 시점에서는 보통 해시를 데이터베이스에 저장합니다.

이제 새 입력이 해시와 같은 데이터인지 확인해야 할 때는 비교 함수를 사용하면 됩니다.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

완료된 해시를 로그한 후 비교 함수를 호출하기 전에 기다려야 하므로, 기존 해시 함수에 이 코드를 추가하고 비교 함수 내에서 'res'를 콘솔에 출력하세요. 콘솔에 해시가 출력되고 이어서 'true'가 출력되는 것을 볼 수 있습니다! 만약 비교 함수에서 'myPlaintextPassword'를 'someOtherPlaintextPassword'로 바꾸면 false가 출력됩니다.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

정확히 작성했다고 생각되면 페이지를 제출하세요.

# --hints--

비동기 해시가 생성되고 올바르게 비교되어야 합니다.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
