---
id: 58a25bcff9fc0f352b528e7e
title: 동기식으로 비밀번호 해시 및 비교하기
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

참고로, 이 프로젝트는 <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>에서 복제한 다음 보일러플레이트 프로젝트를 기반으로 하고 있습니다.

동기식 해시는 구현하기 쉽지만, 서버 측에서 비용이 높거나 해시를 자주 수행할 경우 지연이 발생할 수 있습니다. 이 방법으로 해시하는 것은 다음을 호출하는 것만큼 간단합니다.

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

이 해시 방법을 코드에 추가한 뒤 결과를 콘솔에 출력하세요. 변수들은 이미 서버에 정의되어 있어 조정할 필요가 없습니다. 비동기 함수에서와 같은 비밀번호를 해시하더라도 콘솔에 출력되는 결과가 다를 수 있는데, 이는 해시 문자열 세 번째 부분의 처음 22자가 매번 무작위로 생성되는 솔트 때문입니다. 이제 새 동기 해시와 비밀번호 입력을 비교하려면 compareSync 메서드를 사용하세요:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

결과는 불리언 true 또는 false가 됩니다.

# --instructions--

함수를 추가하고 결과를 콘솔에 출력하여 작동하는지 확인하세요.

정확하다고 생각되면 페이지를 제출하세요.

# --hints--

동기식 해시가 생성되고 올바르게 비교되어야 합니다.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```
