---
id: 683ecb010f4c6bafd428802c
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utahakiki seti ya data ya matibabu ili kuhakikisha inazingatia seti ya sheria.

Tangaza kigezo kinachoitwa `medical_records` na ukae na orodha tupu. Katika hatua zinazofuata, utaitumia kuhifadhi data yako ya matibabu.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Unapaswa kuweka thamani ya orodha tupu kwa kigezo chako cha `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
