---
id: 6842ba0086201089ea268c92
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Ukifuata muundo uleule uliotumia katika hatua zilizopita, orodha ya `medical_records` imejazwa kwa ajili yako na data za wagonjwa wengine. Tafadhali jisikie huru kuizuru.

Ifuatayo utaanza kuandika kitendakazi cha kuthibitisha seti ya data. Unda kitendakazi kinachoitwa `validate` chenye kigezo kimoja `data`.

Unataka kuhakikisha kuwa data zako ni orodha au tuple. Kwa hiyo, ndani ya kitendakazi cha `validate`, tangaza kigezo kinachoitwa `is_sequence` na weka thamani ya wito wa `isinstance`. Pitisha `data` kama hoja ya kwanza na tuple yenye `list` na `tuple` kama hoja ya pili.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

Kitendakazi chako cha `validate` kinapaswa kuwa na kigezo kimoja `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

Unapaswa kuwa na kigezo kinachoitwa `is_sequence` ndani ya kitendakazi chako cha `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Unapaswa kuweka `isinstance(data, (list, tuple))` kwenye kigezo chako cha `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
