---
id: 6842ff2b3c45aa4a89785ecb
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Mara tu baada ya tamko lako la `if`, tangaza kigezo kinachoitwa `is_invalid` na ukaweke thamani `False`. Baadaye, utakitumia kama bendera kuendesha tamko la masharti.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `is_invalid` katika kitendakazi chako.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_invalid")
`) })
```

Unapaswa kuweka thamani `False` kwa kigezo chako cha `is_invalid`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_variable("is_invalid").is_equivalent("is_invalid = False")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False

--fcc-editable-region--
    
--fcc-editable-region--
```
