---
id: 6842d659a4ae7973dabab3b0
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Kamilisha kamusi kwa kuongeza jozi za ufunguo-thamani tatu zifuatazo:

- Ufunguo `diagnosis` wenye thamani ya `Hypertension`
- Ufunguo `medications` wenye thamani ya orodha `['Lisinopril']`
- Ufunguo `last_visit_id` wenye thamani ya `V2301`.

# --hints--

Unapaswa kuwa na ufunguo `diagnosis` wenye thamani ya mfuatano wa herufi `Hypertension` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

Unapaswa kuwa na ufunguo `medications` wenye thamani ya orodha `['Lisinopril']` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

Unapaswa kuwa na ufunguo `last_visit_id` wenye thamani ya mfuatano wa herufi `V2301` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

Unapaswa kuwa na ufunguo `patient_id` wenye thamani ya mfuatano wa herufi `P1001` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

Unapaswa kuwa na ufunguo `age` wenye thamani ya nambari kamili `34` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Unapaswa kuwa na ufunguo `gender` wenye thamani ya mfuatano wa herufi `Female` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
