---
id: 6842d5b537c64266420faf62
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Ongeza ufunguo `age` wenye thamani ya nambari kamili `34` na ufunguo `gender` wenye thamani ya mfuatano wa herufi `Female` katika kamusi yako. Usisahau koma kati ya jozi za ufunguo-thamani.

# --hints--

Orodha yako ya `medical_records` inapaswa kuwa na kamusi moja tu.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Unapaswa kuwa na ufunguo wa `age` wenye thamani ya nambari kamili `34` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Unapaswa kuwa na ufunguo wa `gender` wenye thamani ya mfuatano wa herufi `Female` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

Unapaswa kuwa na ufunguo wa `patient_id` wenye thamani ya mfuatano wa herufi `P1001` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
