---
id: 6842aee5bcaf4ab8883f8191
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Kama ulivyojifunza katika mafundisho yaliyopita, kamusi ni muundo wa data unaohifadhi jozi za ufunguo-thamani, ambapo funguo ni aina za data zisizobadilika (kama mfuatano wa herufi) na ni za kipekee ndani ya kamusi:

```py
person = {
    'name': 'John',
    'age': 33
}
```

Orodha ya `medical_records` itahifadhi kamusi, kila moja ikiwakilisha mgonjwa. Ongeza kamusi yenye funguo `patient_id` na thamani ya mfuatano wa herufi `P1001` kwenye orodha ya `medical_records`.

# --hints--

Orodha yako ya `medical_records` inapaswa kuwa na kamusi moja tu.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Unapaswa kuwa na funguo ya `patient_id` yenye thamani ya mfuatano wa herufi `P1001` katika kamusi yako.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
