---
id: 6846a1f2805ac73ce671a5d8
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Kama ulivyojifunza katika funzo lililopita, kitendakazi cha `enumerate` hukuwezesha kufuatilia kielezo wakati unaporudia juu ya inayoweza kurudiwa:

```py
person = {'name': 'John', 'age': 33}

for index, item in enumerate(person):
    print(index, item)

# 0 name
# 1 age
```

Tengeneza mzunguko wa `for` unaorudia juu ya `data`. Tumia kitendakazi cha `enumerate` kupata kielezo na kipengee katika `data` kila kurudia. Tumia `index` na `dictionary` kama vigezo vya kurudia.

Kwa sasa tumia `pass` kujaza mwili wa mzunguko.

# --hints--

Unapaswa kuwa na mzunguko wa `for` katika kitendakazi chako cha `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0]
`) })
```

Mzunguko wako wa `for` unapaswa kurudia juu ya `enumerate(data)`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_iter().is_equivalent("enumerate(data)")
`) })
```

Unapaswa kutumia `index` na `dictionary` kama vigezo vya kurudia katika mzunguko wako wa `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_for_vars().is_equivalent("index, dictionary")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

--fcc-editable-region--
    
--fcc-editable-region--
```
