---
id: 6849311be42a0431b8e59109
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Unakagua kama data iliyotumwa kwa kitendakazi chako ni orodha au tuple. Bado unahitaji kuhakikisha kuwa kila kipengee katika mfululizo ni kamusi.

Ndani ya mzunguko wako wa `for`, kama kipengee katika `dictionary` si mfano wa `dict`, chapisha `Invalid format: expected a dictionary at position <index>.` (ambapo `<index>` inapaswa kubadilishwa na kielezo cha sasa) na weka `is_invalid` kuwa `True`.

# --hints--

Unapaswa kuwa na tamko la `if` ndani ya mzunguko wako wa `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Tamko lako la `if` linapaswa kukagua kama `dictionary` si mfano wa `dict`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(dictionary, dict)")
`) })
```

Unapaswa kuchapisha `Invalid format: expected a dictionary at position <index>.` (ambapo `<index>` inapaswa kubadilishwa na kielezo cha sasa) ndani ya tamko lako la `if`.

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
validate([{}, "", {}, [], 0])
print = built_in_print

_first = 'Invalid format: expected a dictionary at position 1.'
_second = 'Invalid format: expected a dictionary at position 3.'
_third = 'Invalid format: expected a dictionary at position 4.'
len(out) >= 3 and _first in out and _second in out and _third in out
`)) })
```

Unapaswa kuweka `is_invalid` kuwa `True` ndani ya tamko lako la `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
