---
id: 669fc7e141e4703748c558bf
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do elemento `html`, adicione um elemento `head`.

# --hints--

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de fechamento para `head`.

```js
assert.match(code, /<\/head>/i);
```

A tag de abertura de `head` deve vir antes da tag de fechamento de `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  
--fcc-editable-region--
</html>
```
