---
id: 669fde2081f65141ad703fe4
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

A primeira seção na página será a seção About. A seção apresentará o Sr. Whiskers e dê aos usuários uma ideia do que este blog trata. 

Dentro do elemento `main`, adicione um elemento `section` com o atributo `id` definido como `"about"`. 

Dentro do elemento `section`, adicione um `h2` com o texto de `About`.

# --hints--

O elemento `section` deve ter uma tag de abertura. 

```js
assert.match(code, /<section.*>/);
```

O elemento `section` deve ter uma tag de fechamento. 

```js
assert.match(code, /<\/section\>/);
```

O elemento `section` deve ter um `id` definido como `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement?.getAttribute("id"), "about");
```

O elemento `section` deve estar dentro do elemento `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

O elemento `h2` deve ter uma tag de abertura. 

```js
assert.match(code, /<h2\>/);
```

O elemento `h2` deve ter uma tag de fechamento. 

```js
assert.match(code, /<\/h2\>/);
```

Seu elemento `h2` deve ter o texto `About`. Verifique sua ortografia.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element?.innerText.trim(), "About");
```

O elemento `h2` deve estar dentro do elemento `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
