---
id: 669fdd3965f36f40af9615e5
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

Agora que você terminou de criar o cabeçalho da página, precisará começar a adicionar o conteúdo principal. 

Abaixo do `header`, adicione um elemento `main`. 

# --hints--

O elemento `main` deve ter uma tag de abertura. 

```js
assert.match(code, /<main\>/);
```

O elemento `main` deve ter uma tag de fechamento. 

```js
assert.match(code, /<\/main\>/);
```

O elemento `main` deve estar abaixo do elemento `header`.

```js
assert.match(code, /<\/header>[.\n\s]*<main>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
