---
id: 669fd003cb89ee3c2402e041
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Neste conteúdo introdutório, o objetivo será mostrar uma imagem do sr. Whiskers com uma legenda.

Abaixo do elemento `h1`, comece adicionando um elemento `figure`.

Dentro do elemento `figure`, adicione um elemento `img`.

O atributo `src` do `img` deve ter um valor de `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` e o texto `alt` deve ter um valor de `"a cat in the garden"`.

Abaixo do seu elemento `img`, adicione um `figcaption` com o texto `Mr. Whiskers in the Garden`.

# --hints--

O elemento `figure` deve ter uma tag de abertura. 

```js
assert.match(code, /<figure\>/);
```

O elemento `figure` deve ter uma tag de fechamento. 

```js
assert.match(code, /<\/figure\>/);
```

O elemento `figure` deve estar dentro do elemento `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Você deve ter um elemento `img`.

```js
assert.exists(document.querySelector("img"));
```

O elemento `img` deve estar dentro do elemento `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

O elemento `img` deve ter um atributo `src` com o valor `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Seu elemento `img` deve ter um atributo `alt` definido como `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

O elemento `figcaption` deve ter uma tag de abertura. 

```js
assert.match(code, /<figcaption\>/);
```

O elemento `figcaption` deve ter uma tag de fechamento. 

```js
assert.match(code, /<\/figcaption\>/);
```

O elemento `figcaption` deve estar dentro do elemento `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Seu elemento `figcaption` deve conter o texto `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers in the Garden")
```

O `figcaption` deve estar abaixo do elemento `img`. Eles estão na ordem errada.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```
