---
id: 669fcb06c3034a39f5431a38
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Para preparar para a criação do conteúdo de fato, adicione um elemento `body` abaixo do elemento `head`.

# --hints--

Você deve ter uma tag de abertura `<body>`.

```js
assert.match(code, /<body>/i);
```

Você deve ter uma tag de fechamento `</body>`.

```js
assert.match(code, /<\/body>/i);
```

Você não deve alterar seu elemento `head`. Certifique-se de que não deletou sua tag de fechamento.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

O elemento `body` deve estar depois do elemento `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  
  --fcc-editable-region--
</html>
```
