---
id: 669fcb5106eeb13ab7afcd09
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Neste workshop, você praticará o uso de HTML semântico construindo uma página de blog dedicada ao Sr. Whiskers o gato. 

A primeira seção que você criará é o <dfn>cabeçalho</dfn> da página. 

O elemento `header` é usado para representar o conteúdo introdutório como a navegação da página e outras informações iniciais. 

Aqui está um exemplo usando o elemento `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Dentro do elemento `body`, adicione um elemento `header`.

# --hints--

Você deve ter uma tag de abertura para `header`.

```js
assert.match(code, /<header>/i);
```

Você deve ter uma tag de fechamento para `header`.

```js
assert.match(code, /<\/header>/i);
```

O elemento `header` deve estar dentro do elemento `body`.

```js
assert.exists(document.querySelector('body > header'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
