---
id: 669fdc11c9b0263fe0814a7a
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

Dentro de cada um dos elementos `li`, você precisará ter um elemento âncora.

Para o primeiro elemento anchor, o texto deve ser `About` e o valor do atributo `href` deve ser `"#about"`. O símbolo de hash na frente de `about` representa um nome `id`, que será adicionado mais tarde no projeto.

Para o segundo elemento âncora, o texto deve ser `Posts` e o valor do atributo `href` deve ser `"#posts"`.

Para o terceiro elemento âncora, o texto deve ser `Contact` e o valor do atributo `href` deve ser `"#contact"`.

# --hints--

O primeiro elemento de âncora deve ter o texto `About` dentro do primeiro elemento `li`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "About");
```

O primeiro elemento de âncora deve ter um atributo `href` com o valor de `"#about"`.

```js
const anchorElement = document.querySelector("ul li:first-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#about");
```

O segundo elemento de âncora deve ter o texto `Posts` dentro do segundo elemento `li`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
assert.strictEqual(anchorElement?.innerText.trim(), "Posts");
```

O segundo elemento de âncora deve ter um atributo `href` com o valor de `"#posts"`.

```js
const anchorElement = document.querySelector("ul li:nth-child(2) a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#posts");
```

O terceiro elemento de âncora deve ter o texto `Contact` dentro do terceiro elemento `li`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
assert.strictEqual(anchorElement?.innerText.trim(), "Contact");
```

O terceiro elemento de âncora deve ter um atributo `href` com o valor de `"#contact"`.

```js
const anchorElement = document.querySelector("ul li:last-child a");
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#contact");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
--fcc-editable-region--
          <li></li>
          <li></li>
          <li></li>
--fcc-editable-region--
        </ul>
      </nav>
    </header>
  </body>
</html>
```
