---
id: 67d1edfde28a4133ba64ad26
title: Passo 18
challengeType: 0
dashedName: step-18
---

# --description--

Para melhorar a experiência do usuário, você vai querer aprimorar o número de telefone para que os usuários toquem nele e iniciem uma chamada.

Aqui está como você pode tornar números de telefone clicáveis:

```html
<a href="tel:2345678912">234-567-8912</a>
```

Envolva o texto `555-555-5555` em um elemento anchor e use `tel:` para torná-lo um número de telefone clicável.

# --hints--

Você deve envolver o texto `555-555-5555` em um elemento `a`.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
const link = firstPElement.querySelector('a');
assert.exists(link);
assert.strictEqual(link.textContent.trim(), "555-555-5555");
```

Seu elemento `a` deve ter um atributo `href`.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
const link = firstPElement.querySelector('a');
assert.exists(link.getAttribute('href'));
```

O atributo `href` para o seu novo elemento `a` deve ser definido como `tel:5555555555`.

```js
const firstPElement = document.querySelector('footer #contact address p:nth-of-type(1)');
const link = firstPElement.querySelector('a');
assert.strictEqual(link.getAttribute('href'), "tel:5555555555");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      <section id="posts">
        <h2>Posts</h2>

        <article>
          <h3>Mr. Whiskers' First Day Home</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Bath</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
        <article>
          <h3>Mr. Whiskers' First Birthday Party</h3>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
          <p>
            Lorem ipsum dolor sit amet consectetur adipisicing elit. Quisquam
            quod, voluptates, quae, quos quibusdam dolorum quia nemo repudiandae
            quidem voluptatum quas. Quisquam quod, voluptates, quae, quos
            quibusdam dolorum quia nemo repudiandae quidem voluptatum quas.
          </p>
        </article>
      </section>
    </main>
--fcc-editable-region--
    <footer>
      <section id="contact">
        <h2>Contact</h2>
        <address>
          <p>Phone: 555-555-5555</p>
          <p>Email: fake@email.com</p>
        </address>
      </section>
    </footer>
--fcc-editable-region--
  </body>
</html>
```
