---
id: a77dbc43c33f39daa4429b4f
title: Construye una función de verificación booleana
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

En este laboratorio, construirás una función que verifica si un valor está clasificado como un primitivo booleano.

Booleanos primitivos son `true` y `false`.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función llamada `booWho` que reciba un argumento.
1. Si el argumento recibido es un primitivo booleano, la función debe devolver `true`.
1. Si el argumento es cualquier otro valor, la función debe devolver `false`.

# --hints--

Deberías tener una función `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` debe devolver `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` debe devolver `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` debe devolver `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` debe devolver `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` debe devolver `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` debe devolver `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` debe devolver `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` debe devolver `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` debe devolver `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` debe devolver `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
