---
id: 5cc0bd7a49b71cb96132e54c
title: 构建数字范围生成器
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

在本实验中，你将构建一个 `rangeOfNumbers` 函数，用于生成指定范围内的数字数组。

示例：

- `rangeOfNumbers(3, 9)` 返回 `[3, 4, 5, 6, 7, 8, 9]`
- `rangeOfNumbers(5, 5)` 返回 `[5]`

要求：

- 使用递归（函数必须调用自身）
- 不使用循环或数组方法（`for`、`while`、`forEach`、`map`、`filter`、`reduce`）

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `rangeOfNumbers` 的函数，该函数接受两个参数：`startNum` 和 `endNum`。
1. 该函数应返回一个整数数组，该数组以由 `startNum` 参数表现的数字开始，以由 `endNum` 参数表现的数字结束（包含）。
1. `startNum` 总是小于或相等于 `endNum`。
1. 你的函数必须通过调用自身来使用递归。它不应使用任何循环语法（`for`、`while`，或诸如 `forEach`、`map`、`filter` 或 `reduce` 之类的高阶函数）。
1. 该函数应通过返回仅包含该单个数字的数组来处理 `startNum` 与 `endNum` 相等的基本情形。
1. 对于递归情形，函数应使用修改后的参数调用自身以构建数组，然后将当前数字添加到结果中。
1. 该函数不应依赖全局的变量来缓存或构建数组。

# --hints--

你应该有一个名为 `rangeOfNumbers` 的函数。

```js
assert.isFunction(rangeOfNumbers);
```

`rangeOfNumbers` 函数应接受两个参数。

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

函数应该返回一个数组。

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

不能包含循环语句（`for` 或者 `while` 或者高阶函数比如 `forEach`、`map`、`filter` 或者 `reduce`）。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` 应该使用递归函数（调用自身）来完成这个挑战。

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` 应该返回 `[1, 2, 3, 4, 5]`。

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` 应该返回 `[6, 7, 8, 9]`。

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` 应该返回 `[4]`。

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

`rangeOfNumbers(10, 15)` 应该返回 `[10, 11, 12, 13, 14, 15]`。

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

`rangeOfNumbers(2, 8)` 应该返回 `[2, 3, 4, 5, 6, 7, 8]`。

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

不应使用全局变量来缓存数组。

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
