---
id: 6606b6b7760d0643c3b4eb29
title: ステップ 19
challengeType: 20
dashedName: step-19
---

# --description--

`try` ブロック内のコードが例外を発生させた場合でも、プログラムを継続して実行したいときは、`pass` 文がこれを実現します。

このコードは動作しますが、`except` キーワードの後に例外の種類を指定することが良い習慣とされています。

`ValueError` が発生する可能性があることがわかっているので、`except` キーワードの後にスペースを空けて、その後に `ValueError` を追加してください。

# --hints--

コード内に `except ValueError:` があるはずです。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\n  return row, col\\nexcept ValueError:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except:
                pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
