---
id: 6606a4641ec48b3a9fe8c2fc
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`enumerate` 組み込み関数はイテラブルを引数に取り、繰り返すことができる enumerate オブジェクトを返します。
このオブジェクトはカウント（デフォルトではゼロから始まります）とイテラブルの値を提供します。

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

上記の例のループは、タプル `0, a`、`1, b`、および `2, c` を出力します。

`find_empty_cell` メソッド内で、`pass` を `for` ループに置き換えて、`enumerate()` 関数を使って数独ボードの各行を繰り返してください。現在の行のインデックスには `row` を、現在の行の要素には `contents` を使います。


# --hints--

`pass` を `for` を繰り返す `enumerate(self.board)` ループに置き換えてください。ループ変数には `row` と `contents` を使います。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
