---
id: 6606927d010be4300a4e5330
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

インスタンス化は空のオブジェクトを作成します。しかし、クラスにはメソッドがあり、これは各インスタンスのためのローカル関数のようなものです。クラス内でメソッドは次のように宣言します。

```py
class ClassName:
    def method_name():
        pass
```

`Board` クラスの中で、`pass` を空のメソッド `spam` に置き換えてください。

# --hints--

`spam` クラスの中に `Board` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
