---
id: 66068fb0bfddba2b7977eb60
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このプロジェクトでは、数独パズルのソルバーを作成しながら、クラスとオブジェクトについて学びます。

Pythonでは、クラスはオブジェクトを作成するための設計図です。クラスから作成されたオブジェクトは、そのクラスのインスタンスです。クラスは次の構文で作成できます。

```py
class ClassName:
    pass
```

ここで`class`はクラスを定義するために必要なキーワードで、`ClassName`は慣例に従って*PascalCase*で書かれたクラス名です。

まず`Board`という名前のクラスを作成してください。

# --hints--

`Board`という名前のクラスを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
