---
id: 6606bc4e5535c0484990ccd5
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

次に、指定された数が数独ボードの特定の行に挿入できるかどうかをチェックするメソッドに取り組みます。

`Board` クラス内に、`valid_in_row` という名前のメソッドを作成し、3つのパラメータ `self`、`row`、`num` を与えてください。ここで `self` はクラスのインスタンスを表し、`row` と `num` はそれぞれ行のインデックスとチェックする数を表します。


# --hints--

`valid_in_row` クラス内に `Board` という名前の新しいメソッドを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

`valid_in_row` メソッドは、順に `self`、`row`、`num` の3つのパラメータを持つ必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
