---
id: 66069167b3307b2f4067b22b
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

クラスの新しいインスタンスは、クラス名の後に丸括弧を付ける関数表記を使って作成します。

クラス定義の外で、`Board` クラスのインスタンスを作成し、`gameboard` という名前の変数に代入してください。

# --hints--

`gameboard` クラスの外で `Board` という変数を宣言してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

`gameboard` 変数には `Board()` の値が入っている必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
