---
id: 6606933d6813a8308c962dd1
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

インスタンスメソッドにするには、慣例で`self`と名付けられた特別なパラメータが必要です。このパラメータはクラスのインスタンスへの参照であり、常に最初のパラメータでなければなりません。

`self`メソッドに`spam`パラメータを追加してください。

# --hints--

`spam`メソッドには`self`パラメータが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
