---
id: 660699aabc59c532f2d556e5
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

インスタンスメソッドを呼び出すには、ドット記法を使う必要があります。

```py
instance_name.method_name()
```

ここで `instance_name` はインスタンスまたはオブジェクト、`method_name` は呼び出したいメソッドです。

`spam` オブジェクトの `gameboard` メソッドを呼び出してください。

# --hints--

`spam` オブジェクトの `gameboard` メソッドを `gameboard.spam()` で呼び出す必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
