---
id: 66069e5759b800364707988e
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`__init__` メソッドに戻ると、解くべきパズルを表す追加のパラメータが必要です。

`board` メソッドに `__init__` という名前の2番目のパラメータを追加し、`gameboard` のインスタンス化を修正して、関数呼び出しの引数として渡すように `puzzle` リストを渡してください。

# --hints--

`__init__` メソッドは、順番に `self` と `board` の2つのパラメータを持つ必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

`gameboard` 変数は `Board(puzzle)` の値を持つべきです。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
