---
id: 6606b961ebcf04460f8af76e
title: ステップ 21
challengeType: 20
dashedName: step-21
---

# --description--

`find_empty_cell` メソッドが正しく動作するかどうかを、`gameboard` に対して呼び出し、その結果を出力してテストしてください。

`find_empty_cell` はパラメータを1つ持つように定義されていますが、メソッドを呼び出す際に引数として値を渡してはいけません。`self` が自動的にメソッドを呼び出す対象のオブジェクトとして渡されるためです。

# --hints--

`gameboard.find_empty_cell()` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
