---
id: 6606c14182435d4bab0de2ee
title: ステップ 30
challengeType: 20
dashedName: step-30
---

# --description--

次に、指定された数独ボードの列に数字を挿入できるかどうかを、その列にすでにその数字が存在しないかを確認して判定するメソッドを作成します。

`Board` クラス内に、`valid_in_col` という名前のメソッドを作成し、3つのパラメータ `self`、`col`、`num` を与えてください。
ここで `col` と `num` は、それぞれ列のインデックスとチェックする数字です。

# --hints--

`valid_in_col` クラス内に `Board` という名前の新しいメソッドを作成してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

`valid_in_col` メソッドは、順に `self`、`col`、`num` の3つのパラメータを持つ必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
