---
id: 6606a219f9efbf38ad496f67
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

ドット記法を使ってインスタンス属性にアクセスすることもできます。

`Board` クラスの外で、`gameboard` オブジェクトを初期化した後、`gameboard.board` を使って `board` オブジェクトの `gameboard` 属性にアクセスし、その結果を画面に出力してください。

# --hints--

`gameboard.board` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
