---
id: 660699119472f332798860ad
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

今、`pass` を `print` の呼び出しに置き換え、文字列 `'Spam!'` を渡してください。

# --hints--

`pass` を削除し、`'Spam!'` メソッド内で文字列 `spam` を出力してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
