---
id: 6606bf4561f8794a0d345919
title: ステップ 27
challengeType: 20
dashedName: step-27
---

# --description--

`valid_in_row` メソッドをテストする時が来ました。`valid_in_row` 上で `gameboard` を呼び出してください。`0` と `8` を引数として渡し、その結果を出力してください。

もう一度、メソッドは3つのパラメータで定義されていますが、呼び出し時には2つの引数だけが渡されていることに注意してください。`self` はメソッドが呼び出されるオブジェクトとして自動的に渡されます。

# --hints--

`valid_in_row(0, 8)` 上で `gameboard` を呼び出した結果を出力する必要があります。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
