---
id: 66069f86f58f85371d47123e
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

属性とは、オブジェクトに関連付けられた変数で、通常の変数としてデータをストアするために使われます。

`__init__` メソッドの中で、`board` パラメータ（`Board` クラスのインスタンスを作成するときに渡される）を `self.board` を使ってインスタンス属性 `board` に代入してください。

`self.board` は `Board` クラスのインスタンスの `board` 属性を指します。これは `Board` クラスから作成されたオブジェクトに属する変数です。

# --hints--

`__init__` メソッドの中で `pass` を削除し、`board` パラメータを `self.board` に代入してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
