---
id: 6606a3ccb1eea93a23c066bf
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

これから、数独ボードの空のセルを見つけるメソッドに取り組みます。

`Board` クラス内に、`find_empty_cell` という名前の空のメソッドを作成し、`self` というパラメータを与えてください。

# --hints--

`find_empty_cell` クラスの中に `Board` という名前のメソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

`find_empty_cell` メソッドには `self` というパラメータが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
