---
id: 66069b992c1c5e3451f3deb0
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

インスタンス化は空のオブジェクトを作成します。`__init__` メソッドは、オブジェクトをカスタマイズされた状態でインスタンス化できる特別なメソッドです。クラスが `__init__` メソッドを実装すると、インスタンス化時に自動的に `__init__` が呼び出されます。

`Board` クラスの中で、`spam` メソッドを削除し、`__init__` パラメータを含む `self` メソッドに置き換えてください。

# --hints--

`spam` クラスに `Board` メソッドがあってはいけません。

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

`__init__` クラスに `Board` メソッドを定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

`__init__` メソッドには `self` パラメータが必要です。

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
