---
id: 66b59829dba144ff1351220f
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora es momento de crear un saludo usando la variable `botName`.

In the previous lectures, you learned how to concatenate strings using template literals like this:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Comienza creando una variable llamada `greeting`.

A continuación, usando la sintaxis de literales de plantilla, asigna una cadena que diga `My name is`, seguida de la variable `botName` y terminando con un punto (`.`).

Finalmente, registra la variable `greeting` en la consola.

# --hints--

Debes tener una variable llamada `greeting`.

```js
assert.isNotNull(greeting);
```

Debes asignar una cadena a tu variable `greeting`.

```js
assert.isString(greeting);
```

Debes usar literales de plantilla para concatenar la cadena `My name is` con la variable `botName` seguida de un punto (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Debes tener una declaración de `console` que registre la variable `greeting` en la consola.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
