---
id: 66b59b12a745e10011158f7b
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

El siguiente paso es crear algunas variables más que se utilizarán en futuros mensajes de bot.

Crea una variable llamada `subject` y asígnale el valor de cadena `"JavaScript"`.

Luego, crea una variable llamada `topic` y asígnale el valor de cadena `"strings"`.

# --hints--

Debes tener una variable llamada `subject`.

```js
assert.isNotNull(subject);
```

Debes asignar una cadena a tu variable `subject`.

```js
assert.isString(subject);
```

Debes asignar la cadena `"JavaScript"` a tu variable `subject`.

```js
assert.equal(subject, "JavaScript");
```

Debes tener una variable llamada `topic`.

```js
assert.isNotNull(topic);
```

Debes asignar una cadena a tu variable `topic`.

```js
assert.isString(topic);
```

Debes asignar la cadena `"strings"` a tu variable `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
