---
id: 66b5928b86e907fdfbf34e56
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora es el momento de establecer el nombre del bot.

Cree una variable llamada `botName` y asígnele el valor de cadena `"teacherBot"`.

# --hints--

Debe tener una variable llamada `botName`.

```js
assert.isNotNull(botName);
```

Su variable `botName` debe ser una cadena.

```js
assert.isString(botName);
```

Su variable `botName` debe tener el valor de `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
