---
id: 66b6d482bbb9e12f2e5ee1ae
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

Para esta próxima parte del proyecto, el bot enseñará a trabajar con la propiedad de cadena `length`.

Comienza creando una nueva variable llamada `strLengthIntro`.

Luego usa la sintaxis de literales de plantilla para asignar la cadena `He aquí un ejemplo de uso de la propiedad length en la palabra [sujeto].` a la variable `strLengthIntro`.

Reemplaza `[subject]` con la variable `subject` como lo hiciste antes.

Finalmente, imprime la variable `strLengthIntro` en la consola.

# --hints--

Deberías tener una variable llamada `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

Tu variable `strLengthIntro` debería ser una cadena.

```js
assert.isString(strLengthIntro);
```

Deberías usar la sintaxis de literales de plantilla para asignar la cadena `He aquí un ejemplo de uso de la propiedad length en la palabra [aquí va el sujeto].` a la variable `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Deberías imprimir la variable `strLengthIntro` en la consola.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
