---
id: 66b6e39031393e30f2c48d0f
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Para obtener la longitud de una cadena, puedes utilizar la propiedad <dfn>length</dfn>. Esta propiedad devuelve el número de caracteres en una cadena.

Aquí hay un ejemplo:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Empieza utilizando la propiedad `length` para obtener la longitud de la cadena `subject` y registra ese valor en la consola.

# --hints--


Debes usar la propiedad `length` para obtener la longitud de la cadena `subject`.

```js
assert.match(code, /subject\.length/);
```

Debes registrar la longitud de la cadena `subject` en la consola.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
