---
id: 66b6e80d6c3f0b329c360283
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

La siguiente parte de este taller es revisar el acceso a caracteres de una cadena.

Comienza mostrando el mensaje `Here is an example of accessing the first letter in the word [subject].` en la consola.

Recuerda reemplazar `[subject]` con la variable `subject` y usar la sintaxis adecuada de literales de plantilla como hiciste en los pasos anteriores.

# --hints--

Tu declaración de `console` debería mostrar el mensaje `Here is an example of accessing the first letter in the word [subject].`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
