---
id: 66b6fdb76441c738719039fa
title: Paso 13
challengeType: 1
dashedName: step-13
---

# --description--

In the lectures, you learned how to access the last character in a string like this:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` siempre te dará el último número de índice para una cadena.

Crea una nueva variable llamada `lastCharacter` y asígnale el valor del último carácter en la variable `subject`.

Luego, muestra el valor de la variable `lastCharacter` en la consola.

# --hints--

Debes tener una variable llamada `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Debes asignar el valor del último carácter de la variable `subject` a la variable `lastCharacter`. Consulta el ejemplo si necesitas ayuda.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Debes mostrar el valor de la variable `lastCharacter` en la `consola`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
