---
id: 66b7137348cfb53fd3ec6c73
title: Paso 17
challengeType: 1
dashedName: step-17
---

# --description--

A continuación, agregará una nueva declaración de `console` que muestra el resultado de usar el método `indexOf` en la variable `learningIsFunSentence` para encontrar la posición de la subcadena `"fun"`.

Debajo de esa declaración de `console`, agregue una nueva declaración de `console` que muestra el resultado de usar el método `indexOf` para encontrar la posición de la subcadena `"learning"`.

Tenga en cuenta lo que muestra la última declaración de `console`.

# --hints--

Su `console` debe mostrar la posición del índice de la subcadena `"fun"` de la variable `learningIsFunSentence`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

Su segunda declaración de `console` debe mostrar la posición del índice de la subcadena `"learning"` de la variable `learningIsFunSentence`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
