---
id: 5d822fd413a79914d39e98ca
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Add opening and closing `html` tags below the `DOCTYPE` so you have a place to start putting some code. Be sure to set the language to English.

# --hints--

A declaração `DOCTYPE` deve estar no ínicio do HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

O elemento `html` deve ter uma tag de abertura e um atributo `lang` com o valor de `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

O elemento `html` deve ter uma tag de fechamento.

```js
assert(code.match(/<\/html\s*>/));
```

As tags `html` devem estar na ordem correta.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Você deve ter apenas um elemento `html`.

```js
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--

--fcc-editable-region--

```
