---
id: 65824561f3478e5371a33ae5
title: Passo 32
challengeType: 20
dashedName: step-32
---

# --description--

O próximo passo é definir a função principal, que será o ponto de entrada do programa interativo de controle de despesas.

Defina uma função chamada `main` sem parâmetros. Preencha o corpo da função com a lista `expenses` que você criou no início deste projeto. Você usará esta lista para armazenar os registros de despesas.

# --hints--

Você deve definir uma função chamada `main()` sem parâmetros.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Você deve mover a lista `expenses` para dentro do corpo da função `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
