---
id: 65822773f4c60a479ec89f66
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Dentro da função `print_expenses`, crie um loop `for` que itere sobre cada item na lista `expenses`. Use `expense` como a variável do loop e mova `pass` para dentro do corpo do loop.

# --hints--

Você deve criar um loop `for` que itere sobre `expenses` na sua função `print_expenses`. Lembre-se de usar `pass` dentro do corpo do loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Você deve usar `expense` como a variável do loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

Você deve ter `pass` apenas dentro do corpo do seu loop.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
