---
id: 6582324b5d37fe4b622bc3c0
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Funções Lambda são funções breves e anônimas em Python, ideais para tarefas simples e únicas.Eles são definidos pela palavra-chave `lambda` e usam a seguinte sintaxe:

```py
lambda x: expr
```

No exemplo acima, `x` representa um parâmetro a ser usado na expressão `expr` e ele age como qualquer parâmetro em uma função tradicional. `expr` é a expressão que é avaliada e retornada quando a função lambda é chamada.

Crie uma variável chamada `test` e atribua a ela uma função lambda que recebe um parâmetro `x` e retorna `x * 2`.

# --hints--

Você deve declarar uma variável chamada `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Você deve atribuir `lambda x: x * 2` à sua variável `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
