---
id: 65823cfc74aa564ffc460489
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

Agora, chame `map()` passando sua função `lambda` como o primeiro argumento e a lista `expenses` como o segundo argumento.

# --hints--

Você deve chamar a função `map()` dentro da função `total_expenses`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("total_expenses").find_calls("map")) == 1`)) })
```

Você deve passar sua função `lambda` como o primeiro argumento para a chamada `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[0].is_equivalent("lambda expense: expense['amount']")`)) })
```

Você deve passar a lista `expenses` como o segundo argumento para a chamada `map()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").find_calls("map")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    lambda expense: expense['amount']
--fcc-editable-region--

expenses = []
```
