---
id: 65823040bdb32949d4c1a96f
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Você precisará de uma função para calcular o valor total das despesas.

Defina uma função chamada `total_expenses` que recebe um parâmetro `expenses`.Preencha o corpo da função com uma instrução `pass` por enquanto.

# --hints--

Você deve definir uma função chamada `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Sua função `total_expenses` deve receber um único parâmetro `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
