---
id: 65823b0167fd894f4a7ea60a
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Em seguida, você vai implementar a mesma lógica dentro da função `total_expenses`.

Por enquanto, exclua tanto a função `test` quanto a chamada `print()`.

# --hints--

Você deve deletar sua variável `test` e a chamada relacionada `print()`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_variable("test")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(sum(map(test, [2, 3, 5, 8])))
--fcc-editable-region--

expenses = []
```
