---
id: 658212ba98182f3e855e85f9
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Neste projeto, você vai construir um rastreador de despesas simples, mas funcional em Python.

Comece definindo uma função chamada `add_expense` que recebe três parâmetros: `expenses`, `amount` e `category`. Use a palavra-chave `pass` para preencher o corpo da função.

# --hints--

Você deve usar a palavra-chave `def` para declarar uma função chamada `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Sua função `add_expense` deve receber três parâmetros: `expenses`, `amount` e `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
