---
id: 65823dde36392f505a39f7c7
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Finalmente, passe sua chamada `map()` para a função `sum()` para obter as despesas totais e retornar o resultado.

# --hints--

Você deve chamar a função `sum()` passando sua chamada atual de `map()` como argumento e retornar o resultado da sua função `total_expenses`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("total_expenses").has_return("sum(map(lambda expense: expense['amount'], expenses))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
--fcc-editable-region--
def total_expenses(expenses):
    map(lambda expense: expense['amount'], expenses)
--fcc-editable-region--

expenses = []
```
