---
id: 65823ff0d4b991510fade1a8
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

Dentro da função `filter_expenses_by_category`, substitua `pass` por uma função `lambda`. Use `expense` como parâmetro e avalie a comparação entre o valor da chave `'category'` do dicionário `expense` e `category` usando o operador de igualdade.

# --hints--

Você deve criar uma função `lambda` que use um parâmetro chamado `expense` e avalie a expressão `expense['category'] == category` dentro da função `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

Você não deve ter `pass` dentro da função `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
