---
id: 65823634a308cb4d6bcc5fa6
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Funções lambda podem ser valiosamente combinadas com a função `map()`, que executa uma função especificada para cada elemento em uma coleção de objetos, como uma lista:

```py
map(lambda x: x * 2, [1, 2, 3])
```

A função a ser executada é passada como o primeiro argumento e o iterável é passado como o segundo argumento.

O resultado do exemplo acima seria `[2, 4, 6]`, onde cada item na lista passada para `map()` foi dobrado pela ação da função lambda.

Modifique sua chamada `print()` para imprimir o resultado de chamar `map()` com `test` como o primeiro argumento e `[2, 3, 5, 8]` como o segundo argumento.Você ainda não poderá ver uma saída legível. 

# --hints--

Você deve ter `print(map(test, [2, 3, 5, 8]))` no seu código.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
