---
id: 6582575b8089f85b8b92d7c8
title: Passo 42
challengeType: 20
dashedName: step-42
---

# --description--

Depois de obter os detalhes da despesa, você precisa chamar a função `add_expense` para adicionar a nova despesa à lista `expenses`.

Depois de obter o `amount` e o `category` usando `input()`, chame a função `add_expense`, passando três argumentos: `expenses`, `amount` e `category`. 

- `expenses` é a lista vazia criada na função main anteriormente neste projeto.
- `amount` é o valor da despesa.
- `category` é a categoria da despesa.

# --hints--

Você deve chamar a função `add_expense()` dentro do seu `if` statement.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Você deve passar três argumentos para sua chamada `add_expense()`: `expenses`, `amount` e `category`. A ordem importa.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Sua chamada `add_expense()` deve vir após a atribuição de `amount` e `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
