---
id: 65821c1c40ccac44adf38e6b
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Crie uma lista vazia chamada `expenses`. Você vai usá-lo para armazenar cada uma das suas despesas.

# --hints--

Você deve declarar uma nova variável chamada `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Você deve atribuir uma lista vazia à variável `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
