---
id: 658225d560369446ece5518b
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Adicione outro par chave-valor ao dicionário que você está adicionando à lista `expense`. Use a string `'category'` como chave e o parâmetro `category` como valor.

# --hints--

Você deve adicionar um segundo par chave-valor ao dicionário que está passando como argumento para o método `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict) and len(node[0].tree.keys) == 2
`)) })
```

Você deve passar `{'amount': amount, 'category': category}` como argumento para o método `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount, 'category': category})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount})
--fcc-editable-region--

expenses = []
```
