---
id: 65821fcc010c3245718f2a06
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

O parâmetro `expenses` da sua função `add_expense` será uma lista de despesas. Você quer poder adicionar itens ao final da sua lista. Para isso você usará o método de lista `.append()`.

Adicione uma chamada ao método `.append()` na lista `expenses`. Não passe nenhum argumento para `.append()` por enquanto.

# --hints--

Você deve ter `expenses.append()` na sua função.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

Você não deve ter `pass` na sua função.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
