---
id: 6582507654b3ed5712341382
title: Passo 38
challengeType: 20
dashedName: step-38
---

# --description--

Você vai usar declarações condicionais para verificar a escolha do usuário.Se a escolha for `'1'`, isso significa que o usuário quer adicionar uma despesa.

Ainda no loop `while`, sob a variável `choice`, escreva uma instrução `if` para verificar se `choice` é igual à string `'1'`.Se for verdadeiro, será o ponto de partida para adicionar uma nova despesa.

Dentro do corpo da declaração `if`, declare uma variável `amount` e atribua a ela uma chamada vazia de `input()`.

# --hints--

Você deve criar uma declaração `if` que verifica se `choice` é igual a `'1'` no seu loop `while`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

Você deve declarar uma variável chamada `amount` no corpo da sua declaração `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

Você deve atribuir uma chamada vazia de `input()` à sua variável `amount`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
