---
id: 65822934c8a0454803e3ef12
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Em seguida, você vai exibir os detalhes de cada despesa.

Dentro do loop `for`, substitua `pass` por uma chamada `print()` e passe para ela a seguinte f-string: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Você deve imprimir `f'Amount: {expense}, Category: {expense}'` no seu loop `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

Você não deve ter `pass` dentro do corpo do seu loop.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
