---
id: 658238f7604f154ea9a23e1e
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

A função `sum()` retorna a soma dos itens no iterável que é passado como seu argumento.Você vai usar `sum()` junto com `map()` e funções `lambda` para obter o valor total das despesas.

Por enquanto, faça um pequeno teste e modifique sua chamada atual de `print()` substituindo a chamada de `list()` por uma chamada da função `sum()` passando a chamada atual de `map()` como argumento.

# --hints--

Você deve imprimir `sum(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*sum\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(list(map(test, [2,3,5,8])))
--fcc-editable-region--

expenses = []
```
