---
id: 6582352cbdbcde4d0c4f7b0b
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Para chamar uma função lambda você pode usar a sintaxe usual de função com um par de parênteses após o nome da variável.

Chame sua função lambda `test` e passe `3` como argumento. Então, imprima o resultado.

# --hints--

Você deve ter `print(test(3))` no seu código.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
