---
id: 65824111a09164518320088d
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

A função `filter()` permite que você selecione itens de um iterável, como uma lista, com base na saída de uma função: 

```py
filter(my_function, my_list)
```

`filter()` recebe uma função como seu primeiro argumento e um iterável como seu segundo argumento.Ele retorna um iterador, que é um objeto especial que permite iterar sobre os elementos de uma coleção, como uma lista.

O resultado do exemplo acima é um iterador contendo os elementos de `my_list` para os quais `my_function` retorna `True`. 

Dentro da função `filter_expenses_by_category`, chame `filter()` passando a função `lambda` que você escreveu na etapa anterior como o primeiro argumento e a lista `expenses` como o segundo argumento.

# --hints--

Você deve chamar `filter()` dentro da função `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Você deve passar `lambda expense: expense['category'] == category` como o primeiro argumento para a chamada `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Você deve passar `expenses` como o segundo argumento para a chamada `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
