---
id: 65822308ded1ba4632f66aa3
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Um dicionário é outro tipo de dado embutido em Python. Um dicionário é uma coleção de dados na forma de pares *key*-*value*. Dicionários são definidos com chaves (`{}`) e contêm pares chave-valor separados por vírgulas. Cada chave é seguida por dois pontos (`:`) e pelo valor:

```py
{'amount': 50.0, 'category': 'Food'}
```

No exemplo acima, `'amount'` e `'category'` são chaves e `50.0` e `'Food'` são seus valores correspondentes.

Crie um dicionário com a chave `'amount'` e o valor do parâmetro `amount` e passe seu novo dicionário para a chamada `.append()`.

# --hints--

Você deve passar um dicionário como argumento para o método `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Você deve passar `{'amount': amount}` como argumento para o método `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
