---
id: 6582267ef39f5047411c1743
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Em seguida, defina uma função chamada `print_expenses` que recebe um parâmetro `expenses`.Esta função será usada posteriormente para exibir cada despesa na sua lista.

Preencha o corpo da sua nova função com uma instrução `pass`.

# --hints--

Você deve definir uma nova função chamada `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Sua função `print_expenses` deve receber um único parâmetro `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
