---
id: 658246c28575d653d1f89b59
title: Passo 33
challengeType: 20
dashedName: step-33
---

# --description--

Um loop `while` é outro tipo de loop que executa uma parte do código enquanto uma condição especificada for `True`.O loop termina quando a condição se torna `False`:

```py
while condition:
    <code>
```

Abaixo da lista `expenses`, crie um loop `while`.Use `True` para a condição e imprima a string `'\nExpense Tracker'` dentro do corpo do loop para mostrar o título do programa.

# --hints--

Você deve criar um loop `while` usando `True` como condição dentro da sua função `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Você deve imprimir `'\nExpense Tracker'` no seu loop `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
