---
id: 65822bd82d708c4895080c35
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Em Python, uma coisa importante a saber é que o mesmo tipo de aspas usado para definir uma string não pode ser usado dentro dela.Por exemplo, a string `'I'm a string!'` não é válida.Para usar a aspa simples dentro dessa string você deve:

- Escape a aspa adicionando uma barra invertida antes dela: `'I\'m a string!'`
- Ou use aspas duplas para definir a string: `"I'm a string!"` (preferido).

Você pode acessar valores em um dicionário através de suas chaves.Você precisa usar a notação de colchetes e incluir a chave entre os colchetes quadrados:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Você está atualmente interpolando o dicionário `expense` na sua f-string.Modifique a expressão f-string para acessar o valor da chave `'amount'` e da chave `'category'` no dicionário `expense`.

# --hints--

Você deve passar `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` para sua chamada `print()`.Lembre-se de usar aspas duplas dentro da sua f-string com aspas simples e vice-versa.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
