---
id: 65823e87c9741750a22085a7
title: Passo 28
challengeType: 20
dashedName: step-28
---

# --description--

Em seguida, defina uma função chamada `filter_expenses_by_category` que recebe dois parâmetros: `expenses` e `category`.Use `pass` para preencher o corpo da função.

# --hints--

Você deve definir uma função chamada `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(filter_expenses_by_category)
  `))
})
```

Sua função `filter_expenses_by_category` deve receber `expenses` e `category` como parâmetros.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(filter_expenses_by_category))
    sig == '(expenses, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
