---
id: 66314f2add0a8b57da5d6d91
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Você deve ver algo como `<map object at 0xd273a8>` impresso no console, que é a representação em string do objeto map retornado por `map()`.

Para obter uma saída legível você precisa transformar o objeto map em uma lista.Faça isso passando a chamada `map()` como argumento para a função `list()`.

# --hints--

Você deve passar a chamada `map()` como argumento para a função `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
