---
id: 66694bfa435cef241f6f9a68
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Antes de começar com o projeto, você vai aprender sobre listas.

Ao contrário de números e strings, uma <dfn>list</dfn> é um tipo de dado que funciona como um contêiner para outros valores:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

A lista é caracterizada pelos colchetes ao redor de todos os valores e uma vírgula entre os valores, como: `["A", "happy", "list"]`.Se a lista não contiver nenhum valor, então é uma lista vazia: `[]`.

Uma lista pode conter tipos de dados diferentes: `[1, "Up", ["Down", "Twice"]]`.Isso inclui todos os tipos de dados possíveis.Também pode conter outra lista!

Crie uma variável chamada `my_list` e atribua a ela uma lista vazia.

# --hints--

Você deve ter uma variável chamada `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

A variável `my_list` deve ter o valor de uma lista vazia.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
