---
id: 666951d2e4b27e3a3c72ca5f
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Você também pode acessar um único elemento para obter seu valor.

Listas são indexadas a partir do zero assim como strings.Isso significa que o primeiro elemento está no índice `0`, o segundo elemento está no índice `1` e assim por diante.

Para acessar um elemento você usa a notação de colchetes.Por exemplo, `example_list[1]` está acessando o elemento no índice `1`, o segundo elemento, de `example_list`.

Imprima o primeiro elemento de `my_list`.

# --hints--

Você deve imprimir o primeiro elemento da lista com `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
