---
id: 6669539c1379793f9cb8917c
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Listas Python são mutáveis o que significa que o valor dos itens da lista pode ser alterado. Você pode alterar o valor de um elemento usando a notação de colchetes.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Isso fará com que `example_list` tenha o valor `[4, 'oh', 6, 7]`.

Altere o primeiro elemento de `my_list` para `0`, depois imprima a lista para verificar o valor.

# --hints--

Não altere a declaração da lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

O primeiro elemento de `my_list` deve ser 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Você deve alterar o primeiro elemento usando `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Você deve imprimir a lista.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Você deve imprimir a lista após alterar o primeiro elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
