---
id: 66694fc4bba24f33ca01fa5b
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Python tem alguns métodos de lista. Como métodos para adicionar ou remover itens da lista.

Você pode adicionar um item ao final de uma lista usando o método `append()`. Por exemplo:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` agora é `[4, 5, 6, 7]`.

Tente usar o método `append()` para adicionar `3` a `my_list`. Então imprima a lista.

# --hints--

O último elemento de `my_list` deve ser `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` deve ter um comprimento de 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Você não deve alterar a declaração da lista.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Você deve imprimir a lista.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Você deve imprimir a lista depois de usar `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
