---
id: 66694eaaeb7d253049fa7883
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Para adicionar itens a uma lista, você pode escrevê-los entre os colchetes, separando os valores com vírgulas.

Adicione os números `1` e `2` dentro da lista.

# --hints--

O primeiro elemento de `my_list` deve ser `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

O segundo elemento de `my_list` deve ser `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` deve ter dois elementos dentro.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Você deve atribuir `[1, 2]` diretamente a `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
