---
id: 666ab82c874add3c95145213
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

O método `pop()` pode ser usado para remover um elemento de uma lista.Por padrão, ele remove o último elemento da lista.Você pode passar um índice como argumento para o método e ele removerá o elemento no índice fornecido.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

Neste caso, `fruits_list.pop(2)` remove o elemento no índice `2` da lista.

Use `pop()` para remover o último elemento de `my_list` e depois imprima `my_list`.

# --hints--

Você deve usar `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` deve ser `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Você deve imprimir `my_list` depois de usar `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
