---
id: 666aae6a5d30a71f1fd7749f
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

O método `insert` pode adicionar um elemento em qualquer posição de uma lista.O primeiro argumento é a posição na qual o elemento deve ser adicionado e o segundo argumento é o elemento a ser adicionado.Por exemplo, aqui está como adicionar um novo elemento na terceira posição de `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Usando `insert()`, adicione `1` em `my_list` na posição correta para que esteja contando para cima, depois imprima a lista.

# --hints--

Você deve adicionar `1` a `my_list` no índice `1` usando `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` deve ter 4 elementos.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` deve ser igual a `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Você deve imprimir a lista após adicionar o novo elemento.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
