---
id: 6690e10ebe2181212abc9652
title: 設計名片
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個帶有 `class` 屬性且值為 `business-card` 的 `div`，該 `div` 將包含所有其他元素。
1. 在 `.business-card` 元素內，應該有一個帶有 `class` 屬性且值為 `profile-image` 的 `img` 元素。你可以將影像來源設定為 `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg`（如果你願意）。應該有一個帶有有意義描述的 `alt`。
1. 在 `.business-card` 元素內，`img` 元素之後，你應該有三個帶有 `class` 屬性的 `p` 元素，分別具有 `full-name`、`designation` 和 `company` 的值。
1. 第一個 `p` 元素應該包含你的名字。
1. 第二個 `p` 元素應該包含你的職稱。
1. 第三個 `p` 元素應該包含你的公司名稱。
1. 第三個 `p` 元素下方應該有一個 `hr` 元素。
1. 在 `hr` 元素之後，應該有兩個 `p` 元素。在第一個 `p` 元素中，`text` 應該是一個電子郵件位址。在第二個 `p` 元素中，`text` 應該是一個電話號碼。
1. 在兩個 `p` 元素之後，應該有一個 `a` 元素，該元素的類別為 `portfolio-link`，文字為 `Portfolio`，並且它應該連結到一個有效的 URL。
1. 在包含作品集連結的 `a` 元素之後應該有一個 `hr` 元素。
1. 你應該有另一個帶有 `class` 屬性且值為 `social-media` 的 `div` 元素。在此元素內，應該有一個 `h2` 元素，其文字為 `Connect with me`。
1. 在 `.social-media` 元素內，應該有三個 `a` 元素。這些 `a` 元素應分別有文字 `Twitter`、`LinkedIn` 和 `GitHub`，並且聯結到有效的網站。
1. 你應該在 HTML 檔案的 `head` 區段中使用 `link` 頁籤來連結樣式表 `styles.css`。
1. 你的頁面背景應該是 `rosybrown`。整體字型應該是 `Arial`，備用字型為 `sans-serif`。
1. `.business-card` 選擇器應該具有將元素設定為寬度為 `300px` 且背景顏色為你選擇的屬性。你也應該設定四周的 `padding` 為 `20px`，以及頂部的 `margin` 為 `100px`。文字應置中對齊，字型大小應為 `16px`。左右的邊距應設定為 `auto`。
1. `.profile-image` 選擇器應該有一個 `max-width` 屬性，值為 `100%`。
1. 所有的 `p` 元素應該有上下各 `5px` 的 `margin`。
1. 頁面上的所有聯結都不應該有底線。

# --hints--

你應該有一個主要的 `div`，其 `class` 屬性值為 `business-card`。

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

在 `.business-card` 元素內，應該有一個帶有 `class` 屬性且值為 `profile-image` 的 `img` 元素。

```js
assert.exists(document.querySelector('img.profile-image'));
```

影像原始碼應設定為有效的影像。

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

影像的 `alt` 屬性應設定為有意義的描述。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

在 `.business-card` 元素內，`img` 元素之後，應該有一個帶有 `class` 屬性且值為 `full-name` 的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

在 `.full-name` 元素之後，應該有一個帶有 `class` 屬性且值為 `designation` 的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

在 `.designation` 元素之後，應該有一個帶有 `class` 屬性且值為 `company` 的 `p` 元素。

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

第一個 `p` 元素，也就是具有 `full-name` 類別的那個，應該包含你的名字。

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

第二個 `p` 元素，也就是具有 `class` 屬性且值為 `designation` 的元素，應該包含你的職稱。

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

第三個 `p` 元素，即具有值為 `company` 的 `class` 屬性的元素，應該包含你的公司名稱。

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

在第三個 `p` 元素之後，也就是具有值為 `company` 的 `class` 屬性的那個，應該有一個 `hr` 元素。

```js
assert.exists(document.querySelector('p.company + hr'));
```

在第一個 `hr` 元素之後，應該有一個 `p` 元素，其文字為電子郵件位址。

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

在電子郵件的 `p` 元素之後，應該有另一個 `p` 元素，其文字為你的電話號碼。

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

在電話號碼 `p` 元素之後，應該有一個帶有 `portfolio-link` 類別的 `a` 元素。 

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

`.portfolio-link` 元素應該有 `Portfolio` 的文字。

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

你的 `.portfolio-link` 元素應該有有效的 `href`。

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

在包含你的作品集的 `a` 元素之後，應該有另一個 `hr` 元素作為分隔線。

```js
assert.exists(document.querySelector('a + hr'));
```

在第二個 `hr` 元素之後，應該有一個帶有 `class` 屬性且值為 `social-media` 的 `div` 元素。

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

在 `.social-media` 元素內，應該有一個帶有文字 `Connect with me` 的 `h2` 元素。

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

在 `.social-media` 元素內應該有三個 `a` 元素。

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

第一個 `a` 元素應該有文字 `Twitter`。

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

第一個 `a` 元素的 `href` 應該指向有效的聯結。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

第二個 `a` 元素應該有文字 `LinkedIn`。

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

第二個 `a` 元素的 `href` 應指向有效的聯結。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

第三個 `a` 元素應該有文字 `GitHub`。

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

第三個 `a` 元素的 `href` 應指向有效的聯結。

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

你應該在 `body` 選擇器中設定頁面背景色為 `rosybrown`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

你應該在 `body` 元素中將頁面字型設定為 `Arial`，並以 `sans-serif` 作為備用字型。

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

你的 `.business-card` 選擇器應該有一個 `width` 屬性，值為 `300px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

你的 `.business-card` 選擇器應該有一個帶有有效顏色值的 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

你的 `.business-card` 選擇器應該有一個 `padding` 屬性，值為 `20px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

你的 `.business-card` 選擇器應該有 `100px` 的上邊距。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

`.business-card` 元素內的文字應置中對齊。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

`.business-card` 元素內的文字字型大小應為 `16px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

`.business-card` 元素的左右邊距應設為 `auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

你的 `.profile-image` 選擇器應該有一個 `max-width` 屬性，值為 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

你的 `p` 元素應該有上下各 `5px` 的 `margin`。

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

頁面的聯結不應該有底線。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
