---
id: 6823f9df49cc206af5471a30
title: Eseguire il debug della pagina profilo di Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot sta cercando di creare una pagina profilo. Ha chiesto a un amico di controllare il suo codice e gli ha detto che ci sono alcuni errori.

Il tuo compito è correggere tutti gli errori di Camperbot così potrà continuare a costruire la sua pagina profilo. Completa gli elementi nelle storie utente qui sotto e clicca su "Esegui i test" per vedere se hai risolto tutti gli errori.

**Storie utente:**

1. Camperbot sta cercando di usare un elemento `heading2`, ma quell’elemento non esiste. Correggi quei tag in modo che usino un elemento di intestazione di secondo livello corretto.
2. Camperbot sta cercando di aggiungere due paragrafi con `pp`, ma neanche quelli esistono. Correggili in modo che usino tag di paragrafo corretti.
3. Camperbot sta usando un elemento `h3` per la sottotitolo `Background and Interests` ma ha un errore di sintassi. Individua il problema e risolvilo.

# --hints--

Dovresti avere esattamente un elemento `h1` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Non dovresti rimuovere l’elemento `<h1>Hello from Camperbot!</h1>` dalla pagina.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Dovresti avere esattamente un elemento `h2` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Dovresti avere un elemento `h2` con il testo `About`. Ecco un esempio: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Dovresti avere un elemento paragrafo con il testo `My name is Camperbot and I love learning new things.` Ecco un esempio: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Dovresti avere un elemento paragrafo con il testo `I enjoy solving puzzles.` Ecco un esempio: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Dovresti avere esattamente due elementi paragrafo nella pagina.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Dovresti avere esattamente un elemento `h3` nella pagina.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Dovresti avere un elemento `h3` con il testo `Background and Interests`. Ecco un esempio: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Dovresti avere esattamente cinque elementi totali nella pagina.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
