---
id: 5dc174fcf86c76b9248c6eb2
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Gli elementi HTML hanno un tag di apertura come `<h1>` e un tag di chiusura come `</h1>`. Il testo di un elemento va tra i tag di apertura e di chiusura.

Cambia il testo dell'elemento `h1` sottostante da `Hello World` a `CatPhotoApp` e osserva il cambiamento nell'anteprima del browser.

Quando hai finito, premi il pulsante "Verifica il tuo codice" per vedere se è corretto.

# --hints--

Il testo `CatPhotoApp` dovrebbe essere presente nel codice. Controlla la tua ortografia.

```js
assert(code.match(/catphotoapp/i));
```

L'elemento `h1` dovrebbe avere un tag di apertura. Opening tags have this syntax: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

L'elemento `h1` dovrebbe avere un tag di chiusura. Closing tags have a `/` just after the `<` character.

```js
assert(code.match(/<\/h1\>/));
```

Hai più di un elemento `h1`. Rimuovi l'elemento `h1` di troppo.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Il testo dell'elemento `h1` dovrebbe essere `CatPhotoApp`. Hai omesso del testo, hai un refuso o il testo non è tra i tag di apertura e chiusura dell'elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

