---
id: 5dc1798ff86c76b9248c6eb3
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Gli elementi di intestazione `h1` fino a `h6` sono usati per indicare l'importanza del contenuto che segue. Più basso è il numero, maggiore è l'importanza, quindi gli elementi `h2` hanno meno importanza rispetto agli elementi `h1`.

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Usa solo un elemento `h1` per pagina e posiziona le intestazioni di importanza inferiore sotto quelle di importanza superiore.

Sotto l'elemento `h1`, aggiungi un elemento `h2` con questo testo:

`Cat Photos`

# --hints--

Il tuo elemento `h1` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/h1\>/));
```

Dovresti avere solo un elemento `h1`. Rimuovi quelli in eccesso.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Il testo del tuo elemento `h1` dovrebbe essere `CatPhotoApp`. Hai omesso il testo oppure c'è un errore di battitura.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

I tuoi tag `h2` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML sono scritti in minuscolo.

```js
assert.notMatch(code, /<\/?H2>/);
```

Il tuo elemento `h2` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Il tuo elemento `h2` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/h2\>/));
```

Il testo del tuo elemento `h2` dovrebbe essere `Cat Photos`. Inserisci solo il testo `Cat Photos` tra i tag di apertura e chiusura `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Il tuo elemento `h2` dovrebbe essere sotto l'elemento `h1`. L'elemento `h1` ha maggiore importanza e deve essere sopra l'elemento `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
