---
id: 5dc174fcf86c76b9248c6eb2
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Gli elementi HTML hanno tag di apertura come `<h1>` e tag di chiusura come `</h1>`. Il testo che un elemento mostra va tra il suo tag di apertura e il tag di chiusura.

Cambia il testo dell'elemento `h1` qui sotto da `Hello World` a `CatPhotoApp` e osserva la modifica nell'anteprima del browser.

Quando hai finito, premi il pulsante "Controlla il tuo codice" per vedere se è corretto.

# --hints--

Il testo `CatPhotoApp` dovrebbe essere presente nel codice. Potresti voler controllare l'ortografia.

```js
assert(code.match(/catphotoapp/i));
```

Il tuo elemento `h1` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

I tuoi tag `h1` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML si scrivono in minuscolo.

```js
assert.notMatch(code, /<\/?H1>/);
```

Il tuo elemento `h1` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/h1\>/));
```

Hai più di un elemento `h1`. Rimuovi l'elemento `h1` in eccesso.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Il testo del tuo elemento `h1` dovrebbe essere `CatPhotoApp`. Hai omesso il testo, c'è un errore di battitura oppure non si trova tra il tag di apertura e il tag di chiusura dell'elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
