---
id: 5dc17d3bf86c76b9248c6eb4
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

L'elemento `p` si usa per creare un paragrafo di testo sui siti web. Crea un elemento `p` sotto il tuo elemento `h2` e inserisci il seguente testo:

`Everyone loves cute cats online!`

# --hints--

Il tuo elemento `p` dovrebbe avere un tag di apertura. I tag di apertura hanno la seguente sintassi: `<elementName>`.

```js
assert(document.querySelector('p'));
```

I tuoi tag `p` devono essere in minuscolo. Per convenzione, tutti i tag HTML si scrivono in minuscolo.

```js
assert.notMatch(code, /<\/?P>/);
```

Il tuo elemento `p` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un tag di chiusura subito dopo il carattere `<`.

```js
assert(code.match(/<\/p\>/));
```

Il testo del tuo elemento `p` dovrebbe essere `Everyone loves cute cats online!` Hai omesso il testo oppure c’è un errore di battitura.

```js
const extraSpacesRemoved = document
  .querySelector('p')
  .innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

Il tuo elemento `p` dovrebbe essere sotto l’elemento `h2`. Li hai nell’ordine sbagliato.

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H2') < collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>
    
--fcc-editable-region--
  </body>
</html>
```
