---
id: 5dc23991f86c76b9248c6eb8
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Nel passaggio precedente, hai messo gli elementi `h1`, `h2`, commento e `p` dentro l'elemento `main`. Questo si chiama *annidamento*. Gli elementi annidati dovrebbero essere posizionati due spazi più a destra rispetto all'elemento in cui sono annidati. Questa spaziatura si chiama indentazione ed è usata per rendere l'HTML più facile da leggere.

Ecco un esempio di annidamento e indentazione:

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

L'elemento `h1`, l'elemento `h2` e il commento sono indentati di due spazi in più rispetto all'elemento `main` nel codice qui sotto. Usa la barra spaziatrice sulla tastiera per aggiungere due spazi davanti all'elemento `p` in modo che anche lui sia indentato correttamente.

# --hints--

Il tuo codice dovrebbe avere un elemento `h2` con testo `Cat Photos`. Potresti averlo cancellato per errore, mancano sia il tag di apertura sia il tag di chiusura, oppure il testo è cambiato.

```js
assert(
  document.querySelector('h2') &&
    code.match(/<\/h2\>/) &&
    document.querySelector('h2').innerText.toLowerCase() === 'cat photos'
);
```

Non dovresti aggiungere gli elementi `ul` o `li` dell'esempio.

```js
assert(
  !document.querySelector('ul') && !document.querySelector('li')
);
```

Non dovresti modificare l'indentazione sulla riga con il tuo elemento `h2`. Il suo tag di apertura dovrebbe iniziare a 6 spazi dall'inizio della riga. Puoi riavviare il passaggio per ripristinare l'indentazione originale.

```js
assert(code.toLowerCase().match(/<\/h1\>\s*\n\s{6}<h2>/));
```

Il tuo codice dovrebbe avere un commento. Hai rimosso il commento da un passaggio precedente.

```js
assert(code.match(/<!--.*-->/));
```

Il testo del commento dovrebbe essere `TODO: Add link to cat photos`. Non modificare il testo o la spaziatura del commento.

```js
assert(code.match(/<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\.?\s*-->/i));
```

Non dovresti modificare l'indentazione sulla riga con il tuo elemento commento. Il suo tag di apertura dovrebbe iniziare a 6 spazi dall'inizio della riga. Puoi riavviare il passaggio per ripristinare l'indentazione originale.

```js
assert(
  code
    .toLowerCase()
    .match(/<\/h2>\s*\n\s{6}<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/)
);
```

Il tuo codice dovrebbe avere un elemento `p`. Hai rimosso l'elemento `p` da un passaggio precedente.

```js
assert(document.querySelector('p'));
```

Il testo dell'elemento `p` dovrebbe essere `Everyone loves cute cats online!`. Non modificare il testo, la spaziatura o la punteggiatura dell'elemento `p`.

```js
assert.match(
  document
    .querySelector('p')
    .innerText.toLowerCase(),
    /everyone\s+loves\s+cute\s+cats\s+online!$/
);
```

Il tag di apertura `p` dovrebbe avere un'indentazione che corrisponde a quella degli elementi `h2` e commento. Il suo tag di apertura dovrebbe iniziare a 6 spazi dall'inizio della riga.

```js
assert(code.toLowerCase().match(/-->\s*\n\s{6}<p>/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    </main>
--fcc-editable-region--
  </body>
</html>
```
