---
id: 5dc23f9bf86c76b9248c6eba
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Puoi aggiungere immagini al tuo sito web usando l'elemento `img`. Gli elementi `img` hanno un tag di apertura senza tag di chiusura. Un elemento senza tag di chiusura è noto come <dfn>elemento void</dfn>.

Aggiungi un elemento `img` sotto l'elemento `p`. A questo punto, nessuna immagine apparirà nel browser.

# --hints--

Il tuo elemento `img` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('img'));
```

Il tuo elemento `img` non dovrebbe avere un tag di chiusura `</img>`.

```js
assert(!code.match(/<\/img\>/));
```

Dovresti avere solo un elemento `img`. Rimuovi eventuali elementi in eccesso.

```js
assert(document.querySelectorAll('img').length === 1);
```

Il tuo elemento `img` dovrebbe essere sotto l'elemento `p`. Li hai nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('P') < collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
