---
id: 5dc2385ff86c76b9248c6eb7
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

HTML5 ha alcuni elementi che identificano diverse aree di contenuto. Questi elementi rendono il tuo HTML più facile da leggere e aiutano con l'ottimizzazione per i motori di ricerca (SEO) e l'accessibilità.

L'elemento `main` viene usato per rappresentare il contenuto principale del corpo di un documento HTML. Il contenuto all'interno dell'elemento `main` dovrebbe essere unico per il documento e non dovrebbe essere ripetuto in altre parti del documento.

```html
<main>
  <h1>Most important content of the document</h1>
  <p>Some more important content...</p>
</main>
```

Identifica la sezione principale di questa pagina aggiungendo un tag di apertura `<main>` prima dell'elemento `h1` e un tag di chiusura `</main>` dopo l'elemento `p`.

# --hints--

Il tuo elemento `main` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('main'));
```

I tuoi tag `main` dovrebbero essere in minuscolo. Per convenzione, tutti i tag HTML sono scritti in minuscolo.

```js
const mainTagMatches = code.matchAll(/<\/?main>/gi);
for (const match of mainTagMatches) {
  const tag = match[0];
  assert.strictEqual(
    tag,
    tag.toLowerCase()
  );
}
```

Il tuo elemento `main` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/main\>/));
```

Il tag di apertura del tuo elemento `main` dovrebbe essere sotto il tag di apertura dell'elemento `body`. Li hai nell'ordine sbagliato.

```js
const main = document.querySelector('main');
assert.equal(main?.previousElementSibling, null);
```

Il tag di apertura del tuo elemento `main` dovrebbe essere sopra l'elemento `h1`. Li hai nell'ordine sbagliato.

```js
const collection = [...document.querySelectorAll('main,h1')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('MAIN') < collection.indexOf('H1'));
```

Il tag di chiusura del tuo elemento `main` dovrebbe essere sotto l'elemento `p`. Li hai nell'ordine sbagliato.

```js
const mainNode = document.querySelector('main');
const pNode = document.querySelector('p');
assert.isTrue(mainNode.contains(pNode));
assert.match(pNode.textContent.toLowerCase(), /everyone loves cute cats online!/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
  
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
    <!-- TODO: Add link to cat photos -->
    <p>Everyone loves cute cats online!</p>
    
--fcc-editable-region--
  </body>
</html>
```
