---
id: 5dc24073f86c76b9248c6ebb
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Gli <dfn>attributi</dfn> HTML sono parole speciali usate all'interno del tag di apertura di un elemento per controllarne il comportamento. L'attributo `src` in un elemento `img` specifica l'URL dell'immagine (dove si trova l'immagine).

Ecco un esempio di un elemento `img` con un attributo `src` che punta al logo di freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

All'interno dell'elemento `img` esistente, aggiungi un attributo `src` con questo URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Il tuo codice dovrebbe avere un elemento `img`. Potresti aver rimosso l'elemento `img` oppure non aver racchiuso il valore dell'attributo `src` tra virgolette.

```js
assert(document.querySelector('img'));
```

Il tuo elemento `img` dovrebbe avere un attributo `src`. Hai omesso l'attributo oppure c'è un errore di battitura. Assicurati che ci sia uno spazio tra il nome dell'elemento e il nome dell'attributo.

```js
assert(document.querySelector('img').src);
```

L'attributo `img` del tuo elemento `src` dovrebbe essere impostato su '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Hai omesso l'URL oppure c'è un errore di battitura. Il maiuscolo/minuscolo dell'URL è importante.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Anche se hai impostato l'attributo `img` dell'elemento `src` sull'URL corretto, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Il tuo elemento `img` dovrebbe essere chiuso con `>` oppure `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
