---
id: 5ddbd81294d8ddc1510a8e56
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Il testo di un collegamento deve essere inserito tra il tag di apertura e il tag di chiusura di un elemento àncora (`a`).

Ecco un esempio di un collegamento con il testo `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Aggiungi il testo dell'àncora `cat photos` all'elemento àncora. Questo diventerà il testo del collegamento.

# --hints--

Il tuo elemento àncora (`a`) dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Il tuo elemento àncora (`a`) dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/a\>/));
```

Il testo del tuo elemento àncora (`a`) dovrebbe essere `cat photos`. Assicurati di mettere il testo del collegamento tra il tag di apertura e il tag di chiusura dell'elemento àncora (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
