---
id: 5dfa2407b521be39a3de7be1
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Per aprire i collegamenti in una nuova scheda, puoi usare l'attributo `target` sull'elemento àncora (`a`).

L'attributo `target` specifica dove aprire il documento collegato. `target="_blank"` apre il documento collegato in una nuova scheda o finestra.

Ecco la sintassi base per un elemento `a` con un attributo `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Aggiungi un attributo `target` con valore `_blank` al tag di apertura dell'elemento àncora (`a`), così il collegamento si aprirà in una nuova scheda.

# --hints--

Il tuo elemento `p` dovrebbe avere un elemento àncora annidato (`a`) con il testo `cat photos`. Potresti averlo eliminato o avere un errore di battitura.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Il tuo elemento àncora (`a`) non ha un attributo `target`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Il valore dell'attributo `target` dovrebbe essere `_blank`. Hai omesso il valore o c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
