---
id: 5dfa37b9eacea3f48c6300b0
title: Passaggio 22
challengeType: 0
dashedName: step-22
---

# --description--

Per creare una lista non ordinata di elementi, puoi usare l'elemento `ul`. 

Dopo l'elemento `h3` con il testo `Things cats love:`, aggiungi un elemento lista non ordinata (`ul`). Nota che a questo punto non verrà visualizzato nulla.

# --hints--

Il tuo elemento `ul` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

Il tuo elemento `ul` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/ul>/));
```

L'elemento `ul` dovrebbe essere sopra il tag di chiusura del secondo elemento `section`.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
