---
id: 5ef9b03c81a63668521804dd
title: Passaggio 47
challengeType: 0
dashedName: step-47
---

# --description--

Gli elementi `label` servono ad associare il testo di un elemento `input` con l'elemento `input` stesso (soprattutto per tecnologie assistive come i lettori di schermo).

Ecco un esempio di un elemento `label` con un pulsante `radio`:

```html
<label><input type="radio"> cat</label>
```

Nell'esempio, cliccando sulla parola `"cat"` si selezionerà anche il pulsante `radio`.

Annida il tuo pulsante `radio` all'interno di un elemento `label`.

# --hints--

Devi assicurarti che il pulsante di opzione sia ancora presente.

```js
assert(document.querySelector('input[type="radio"]'));
```

Il testo ` Indoor` dovrebbe trovarsi direttamente a destra del tuo pulsante `radio`. Assicurati che ci sia uno spazio tra l'elemento e il testo. Hai omesso il testo oppure c'è un errore di battitura.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

Il tuo elemento `label` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('label'));
```

Il tuo elemento `label` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/label\>/));
```

Il tuo pulsante di opzione e il suo testo dovrebbero trovarsi tutti tra il tag di apertura e il tag di chiusura dell'elemento `label`.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
