---
id: 5ef9b03c81a63668521804d7
title: Passaggio 38
challengeType: 0
dashedName: step-38
---

# --description--

L'attributo `action` indica dove devono essere inviati i dati del modulo.

Ecco un esempio di un elemento `form` con un attributo `action`:

```html
<form action="/submit-url"></form>
```

Nell'esempio, `action="/submit-url"` dice al browser che i dati del modulo devono essere inviati al percorso `/submit-url`.

Aggiungi un attributo `action` con valore `https://freecatphotoapp.com/submit-cat-photo` all'elemento `form`.

# --hints--

Il tuo elemento `form` dovrebbe avere un tag di apertura e un tag di chiusura nel corretto ordine. Potresti mancare uno o entrambi i tag richiesti, oppure averli nell'ordine sbagliato.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Il tuo elemento `form` annidato nell'ultimo elemento `section` dovrebbe essere sotto l'elemento `h2`. Hai l'elemento `h2` e l'elemento `form` nell'ordine sbagliato.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Il tuo elemento `form` non dovrebbe avere contenuto. Rimuovi eventuali elementi HTML o testo tra i tag dell'elemento `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Il tuo elemento `form` non ha un attributo `action`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Il tuo elemento `form` dovrebbe avere un attributo `action` con valore `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Anche se hai impostato l'attributo `action` sull'URL corretto, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
