---
id: 5ef9b03c81a63668521804db
title: Passaggio 43
challengeType: 0
dashedName: step-43
---

# --description--

Per impedire a un utente di inviare il modulo quando mancano informazioni obbligatorie, devi aggiungere l'attributo `required` a un elemento `input`.

Ecco un esempio di un campo di input con l'attributo `required`:

```html
<input type="text" name="firstName" required>
```

Non è necessario impostare un valore per l'attributo `required`. Invece, aggiungi semplicemente la parola `required` all'elemento `input`, assicurandoti che ci sia uno spazio tra essa e gli altri attributi.

# --hints--

Hai eliminato il tuo elemento `input` oppure ha una sintassi non valida. Tutti i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(document.querySelectorAll('input').length);
```

Il tuo `form` dovrebbe contenere solo l'elemento `input`. Rimuovi eventuali altri elementi HTML o testo all'interno dell'elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Il tuo elemento `input` dovrebbe avere un attributo `required`. Ricorda, devi solo aggiungere la parola `required` all'interno del tag dell'elemento `input`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

Non dovrebbe essere assegnato alcun valore all'attributo `required`.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
