---
id: 5ef9b03c81a63668521804d8
title: Passaggio 39
challengeType: 0
dashedName: step-39
---

# --description--

L'elemento `input` ti permette diversi modi per raccogliere dati da un modulo web. Come gli elementi `img`, gli elementi `input` sono un <dfn>elemento vuoto</dfn> e non necessitano di tag di chiusura.

Annida un elemento `input` nell'elemento `form`.

# --hints--

Il tuo elemento `form` dovrebbe avere un tag di apertura e un tag di chiusura nel corretto ordine. Potresti mancare uno o entrambi i tag richiesti, oppure averli nell'ordine sbagliato.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Il tag di apertura del tuo elemento `form` dovrebbe avere solo un attributo `action`. Rimuovi qualsiasi altra cosa tu abbia scritto.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Dovresti creare un elemento `input`. Controlla la sintassi.

```js
assert(document.querySelector('input'));
```

Il tuo elemento `input` dovrebbe avere un tag di apertura, ma non un tag di chiusura.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Il tuo elemento `input` dovrebbe essere annidato all'interno dell'elemento `form`.

```js
assert(document.querySelector('form > input'));
```

Il tuo `form` dovrebbe contenere solo l'elemento `input`. Rimuovi qualsiasi elemento HTML o testo tra i tag dell'elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
