---
id: 5ef9b03c81a63668521804d9
title: Passaggio 42
challengeType: 0
dashedName: step-42
---

# --description--

Il testo segnaposto serve a dare un suggerimento su che tipo di informazione inserire in un input.

Ecco un esempio di un elemento `input` con un segnaposto impostato su `e.g. Jane Doe`:

```html
<input type="text" placeholder="e.g. Jane Doe">
```

Aggiungi il testo segnaposto `cat photo URL` al tuo elemento `input`.

# --hints--

Hai eliminato il tuo elemento `input` oppure la sua sintassi non è valida. Tutti i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(document.querySelectorAll('input').length);
```

Il tuo `form` dovrebbe contenere solo l'elemento `input`. Rimuovi eventuali altri elementi HTML o testo all'interno dell'elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Il tuo elemento `input` non ha un attributo `placeholder`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('placeholder'));
```

Il tuo elemento `input` dovrebbe avere un attributo `placeholder` con valore `cat photo URL`. Hai omesso il valore oppure c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(
  document.querySelector('input')
    ?.getAttribute('placeholder')
    .replace(/\s+/g, ' ')
    .match(/^cat photo URL$/i)
);
```

Anche se hai impostato l'attributo `input` dell'elemento `placeholder` su `cat photo URL`, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+placeholder\s*=\s*cat\s+photo\s+url/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
