---
id: 5ef9b03c81a63668521804e2
title: Passaggio 56
challengeType: 0
dashedName: step-56
---

# --description--

I moduli usano comunemente le caselle di spunta per domande che possono avere più di una risposta. L'elemento `input` con un attributo `type` impostato su `checkbox` crea una casella di spunta.

Sotto l'elemento `legend` che hai appena aggiunto, aggiungi un `input` con il suo attributo `type` impostato su `checkbox` e il suo testo impostato su:

` Loving`

# --hints--

L'elemento `input` per la tua casella di spunta dovrebbe avere un tag di apertura, ma non un tag di chiusura.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Dovresti aver aggiunto solo un elemento input per la tua casella di spunta. Rimuovi eventuali elementi in più.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Il tuo nuovo elemento `input` dovrebbe essere sotto l'elemento `legend` con il testo `What's your cat's personality?`. Li hai nell'ordine sbagliato.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Il tuo nuovo elemento `input` non ha un attributo `type`. Controlla che ci sia uno spazio dopo il nome del tag di apertura.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Il tuo nuovo elemento `input` dovrebbe avere un attributo `type` con il valore `checkbox`. Hai omesso il valore o c’è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Anche se hai impostato l'attributo `input` dell'elemento `type` su `checkbox`, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Il testo ` Loving` dovrebbe trovarsi direttamente a destra della tua casella di spunta. Assicurati che ci sia uno spazio tra l'elemento e il testo. Hai omesso il testo o c’è un errore di battitura.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
