---
id: 5ef9b03c81a63668521804dc
title: Passaggio 46
challengeType: 0
dashedName: step-46
---

# --description--

Puoi usare i pulsanti di opzione per domande in cui vuoi solo una risposta tra più opzioni.

Ecco un esempio di pulsante di opzione con il testo impostato come `cat`:

```html
<input type="radio"> cat
```

Ricorda che un elemento `input` è un elemento vuoto.

Prima dell'input di testo, aggiungi un pulsante di opzione con il testo impostato come:

`Indoor`

# --hints--

Dovresti creare un elemento `input` per il tuo pulsante di opzione. Controlla la sintassi.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Il tuo elemento `input` dovrebbe avere un tag di apertura, ma non un tag di chiusura.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Dovresti aver aggiunto solo un elemento `input` per il tuo pulsante di opzione. Rimuovi eventuali elementi extra.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Il tuo nuovo elemento `input` dovrebbe essere sopra l'esistente `input` con l'attributo `type` impostato su `text`. Li hai nell'ordine sbagliato.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Il tuo nuovo elemento `input` non ha un attributo `type`. Controlla che ci sia uno spazio dopo il nome del tag di apertura.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Il tuo nuovo elemento `input` dovrebbe avere solo un attributo `type`. Rimuovi eventuali extra.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Il tuo nuovo elemento `input` dovrebbe avere un attributo `type` con il valore `radio`. Hai omesso il valore o c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Anche se hai impostato l'attributo `input` del nuovo elemento `type` su `radio`, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

Il testo `radio` del pulsante ` Indoor` dovrebbe essere posizionato dopo di esso invece che prima.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

Il testo ` Indoor` dovrebbe essere posizionato direttamente a destra del tuo pulsante `radio`. Hai omesso il testo o c'è un errore di battitura.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
 
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
