---
id: 5efb23e70dc218d6c85f89b1
title: Passaggio 40
challengeType: 0
dashedName: step-40
---

# --description--

Ci sono molti tipi di input che puoi creare usando l'attributo `type`. Puoi facilmente creare un campo password, un pulsante di reset o un controllo per permettere agli utenti di selezionare un file dal loro computer.

Crea un campo di testo per ricevere un input testuale da un utente aggiungendo l'attributo `type` con il valore `text` all'elemento `input`.

# --hints--

Hai eliminato il tuo elemento `input` oppure ha una sintassi non valida. Se hai aggiunto attributi, assicurati che i loro valori siano racchiusi tra virgolette.

```js
assert(document.querySelectorAll('input').length);
```

Il tuo elemento `form` dovrebbe contenere solo l'elemento `input`. Rimuovi eventuali elementi HTML o testo extra tra i tag dell'elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Il tuo elemento `input` non ha un attributo `type` con il valore `text`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Il tuo elemento `input` dovrebbe avere un attributo `type` con il valore `text`. Hai omesso il valore oppure c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type').replace(/\s+/g, ' '),
  /^text$/i
);
```

Anche se hai impostato l'attributo `input` dell'elemento `type` su `text`, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(!/\<input\s+type\s*=\s*text/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
