---
id: 5efc518e8d6a74d05e68af75
title: Passaggio 59
challengeType: 0
dashedName: step-59
---

# --description--

Aggiungi l'attributo `name` con il valore `personality` all'elemento casella di spunta `input`.

Anche se non noterai nulla nel browser, farlo rende più semplice per un server elaborare il tuo modulo web, soprattutto quando ci sono più caselle di spunta.

# --hints--

Dovresti assicurarti che la casella di spunta sia ancora presente.

```js
assert(document.querySelector('input[type="checkbox"]'));
```

L'elemento casella di spunta `input` non ha un attributo `name`. Controlla che ci sia uno spazio dopo il nome del tag di apertura.

```js
assert.isTrue(document.querySelector('input[type="checkbox"]')?.hasAttribute('name'));
```

L'elemento casella di spunta `input` dovrebbe avere un attributo `name` con il valore `personality`. Hai omesso il valore oppure c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.match(
  document.querySelector('input[type="checkbox"]')?.getAttribute('name'),
  /^personality$/
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
