---
id: 5f05a1d8e233dff4a68508d8
title: Passaggio 49
challengeType: 0
dashedName: step-49
---

# --description--

Crea un altro pulsante di opzione sotto il primo. Annidalo all'interno di un elemento `label` con `Outdoor` come testo `label`. Dai al pulsante di opzione un attributo `id` con `outdoor` come valore.

# --hints--

Solo il pulsante di opzione `Indoor` originale dovrebbe avere un insieme `id` impostato su `indoor`. Puoi riavviare il passaggio per ripristinare l'HTML originale se necessario.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Non dovresti apportare modifiche al pulsante di opzione `Indoor`. Puoi riavviare il passaggio per ripristinare l'HTML originale se necessario.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Dovresti aggiungere esattamente un nuovo elemento `input` annidato in un nuovo elemento `label`. Assicurati che il tuo nuovo `label` abbia sia un tag di apertura sia un tag di chiusura.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Non dovresti aggiungere attributi a nessun tag di apertura `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

Il tuo nuovo `input` dovrebbe avere un attributo `id` con valore `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

Il tuo nuovo `input` dovrebbe avere un attributo `type` con valore `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Non dovresti aggiungere nuovi elementi oltre a un `input` annidato in un `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Non dovrebbe esserci testo a sinistra del tuo nuovo `input`.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Il testo dell'etichetta per il tuo nuovo pulsante di opzione deve essere esattamente `Outdoor`. Questo include la maiuscola.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Il tuo nuovo pulsante di opzione e l'etichetta dovrebbero essere immediatamente sotto/dopo il pulsante di opzione e l'etichetta `Indoor`. Non dovrebbero esserci altri tag tra di loro.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
