---
id: 671141feba228a35cefba82d
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Trasforma il testo esistente `cute cats` in un elemento àncora che colleghi a:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Dovrebbe esserci un nuovo elemento àncora nel primo elemento `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Il testo dell'àncora dovrebbe essere `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

L'attributo `href` del nuovo elemento àncora dovrebbe essere `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Nel tuo codice dovrebbe esserci solo una istanza della frase `cute cats`.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

Il testo dell'elemento `p` dovrebbe essere ancora `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
