---
id: 7cf9b03d81a65668421804c3
title: Passaggio 41
challengeType: 0
dashedName: step-41
---

# --description--

Perché i dati di un modulo possano essere consultati dalla destinazione specificata nell'attributo `action`, devi assegnare al campo di testo un attributo `name` con un valore che rappresenti i dati inviati.

Ecco un esempio di un elemento `input` con un attributo `name`:

```html
<input type="text" name="name">
```

Aggiungi l'attributo `name` con il valore `catphotourl` al tuo campo di testo.

# --hints--

Hai eliminato il tuo elemento `input` oppure la sua sintassi è errata. Tutti i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert(document.querySelectorAll('input').length);
```

Il tuo `form` dovrebbe contenere solo l'elemento `input`. Rimuovi eventuali altri elementi HTML o testo all'interno dell'elemento `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Il tuo elemento `input` non ha un attributo `name`. Controlla che ci sia uno spazio dopo il nome del tag di apertura e/o che ci siano spazi prima di tutti i nomi degli attributi.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Il tuo elemento `input` dovrebbe avere un attributo `name` con il valore `catphotourl`. Hai omesso il valore oppure c'è un errore di battitura. Ricorda che i valori degli attributi devono essere racchiusi tra virgolette.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Anche se hai impostato l'attributo `input` dell'elemento `name` su `catphotourl`, è consigliato racchiudere sempre il valore di un attributo tra virgolette.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
