---
id: 66ea7adae8053065a64f9002
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você revisará como trabalhar com os diferentes métodos do objeto `Math` construindo um Mathbot.Este Mathbot registrará algumas operações matemáticas e mensagens no console na tentativa de ensinar sobre o objeto `Math` do JavaScript.

Para esta primeira etapa, comece criando uma variável chamada `botName` e atribua a ela o valor string `"MathBot"`.

Então, crie outra variável chamada `greeting` e atribua a ela a sentença `"Hi there!Meu nome é [botName goes here] e estou aqui para ensinar você sobre o objeto Math!`No lugar do espaço reservado `[botName goes here]`, use a variável `botName`. 

Você pode usar template literals ou concatenação de strings com o operador `+` para conseguir isso.

Finalmente, registre a variável `greeting` no console.

# --hints--

Você deve ter uma variável chamada `botName`.

```js
assert.isNotNull(botName);
```

A variável `botName` deve ser uma string.

```js
assert.isString(botName);
```

Você deve atribuir o valor string `"MathBot"` à variável `botName`.

```js
assert.strictEqual(botName, "MathBot");
```

Você deve ter uma variável chamada `greeting`.

```js
assert.isNotNull(greeting);
```

Sua variável `greeting` deve ser uma string.

```js
assert.isString(greeting);
```

Você deve atribuir a seguinte frase à variável `greeting`: `"Hi there!Meu nome é [MathBot goes here] e estou aqui para ensinar você sobre o objeto Math!`Certifique-se de substituir `[MathBot goes here]` pela variável `botName`.

```js
assert.strictEqual(greeting, "Hi there! My name is MathBot and I am here to teach you about the Math object!");
```

Você deve registrar a variável `greeting` no console.

```js
assert.match(code, /console\.log\(greeting\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
