---
id: 66ea8239b9eeb76f12299030
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

A próxima parte do workshop é revisar como gerar um número aleatório entre dois valores. 

Comece adicionando outro `console.log()` que registre a mensagem `"Now, generate a random number between two values."` no console.

Então, crie uma variável chamada `min` e atribua a ela o valor `1` e crie uma variável chamada `max` e atribua a ela o valor `100`. 

No próximo passo, você irá gerar um número aleatório entre esses dois valores.

# --hints--

Você deve ter um `console.log()` que registre a mensagem `"Now, generate a random number between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+number\s+between\s+two\s+values\.\s*\1\);?/);
```

Você deve ter uma variável chamada `min`.

```js
assert.isNotNull(min);
```

Sua variável `min` deve ser um número.

```js
assert.isNumber(min);
```

Você deve atribuir o número `1` à sua variável `min`.

```js
assert.equal(min, 1);
```

Você deve ter uma variável chamada `max`.

```js
assert.isNotNull(max);
```

Sua variável `max` deve ser um número.

```js
assert.isNumber(max);
```

Você deve atribuir o número `100` à sua variável `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
