---
id: 66ea845f54195c7020cc521b
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

A fórmula para gerar um número aleatório entre dois valores é a seguinte:

```js
Math.random() * (maximum - minimum) + minimum;
```

Isso produzirá um resultado que é um número de ponto flutuante entre dois valores.

Crie uma variável chamada `randomNum2` e atribua a ela o resultado de gerar um valor entre os valores `min` e `max`.

Então, registre a variável `randomNum2` para ver o resultado.Tente atualizar a página para ver resultados diferentes.

# --hints--

Você deve ter uma variável chamada `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Sua variável `randomNum2` deve ser um número.

```js
assert.isNumber(randomNum2);
```

Você deve atribuir o resultado de gerar um valor entre os valores `min` e `max`.Certifique-se de não codificar um valor fixo.Você deve usar as fórmulas fornecidas acima.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Você deve registrar a variável `randomNum2` no console.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
