---
id: 66ea7fc048a6e26e33c598fe
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora, é hora de gerar um número aleatório usando o método `Math.random()`.

Crie uma variável chamada `randomNum` e atribua a ela o valor do resultado da chamada do método `Math.random()`.

Então, registre a variável `randomNum` no console.

Tente adicionar um espaço no código para reexecutar o bot e ver diferentes números aleatórios que são gerados.

# --hints--

Você deve ter uma variável chamada `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Você deve atribuir o resultado da chamada do método `Math.random()` à variável `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Você deve registrar a variável `randomNum` no console.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
