---
id: 66ea7e7399244c6d4b6052a1
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

A próxima parte do workshop revisará como o método `Math.random()` funciona. 

Comece adicionando outro `console.log()` que registre a string `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."` no console.
 
# --hints--

Você deve ter um `console.log()` que registre a string `"The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1."`.

```js
assert.match(code, /console\.log\((`|'|")\s*The\s+Math\.random\(\)\s+method\s+returns\s+a\s+pseudo\s+random\s+number\s+greater\s+than\s+or\s+equal\s+to\s+0\s+and\s+less\s+than\s+1\.\1\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);
--fcc-editable-region--

--fcc-editable-region--
```
