---
id: 66eb3ec3b07dfe8f898646c0
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Como você aprendeu nas lições anteriores, o método `Math.floor()` arredonda o valor para baixo até o inteiro mais próximo.

```js
const price = 10.99;
Math.floor(price); // 10
```

Crie uma variável chamada `numRoundedDown` e atribua a ela o resultado de arredondar o número de ponto flutuante `6.7` para baixo até o inteiro mais próximo.

Então, registre a variável `numRoundedDown` no console para ver o resultado. 

# --hints--

Você deve ter uma variável `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

A variável `numRoundedDown` deve ser um número.

```js
assert.isNumber(numRoundedDown);
```

Sua variável `numRoundedDown` deve ser igual ao resultado de arredondar o número de ponto flutuante `6.7` para baixo até o inteiro mais próximo.

```js
assert.equal(numRoundedDown, 6);
```

Você não deve codificar o valor `6` diretamente para a variável `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Você deve registrar a variável `numRoundedDown` no console.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
