---
id: 66eb488c0129e192542f0865
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Nas lições anteriores, você aprendeu como trabalhar com o método `Math.ceil()` assim:

```js
const price = 10.01;
Math.ceil(price); // 11
```

O método `Math.ceil()` arredonda o valor para cima até o inteiro mais próximo.

Crie uma variável chamada `numRoundedUp` e atribua a ela o resultado de arredondar o número de ponto flutuante `3.2` para cima até o inteiro mais próximo.

Então, registre a variável `numRoundedUp` no console para ver o resultado.

# --hints--

Você deve ter uma variável `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

A variável `numRoundedUp` deve ser um número.

```js
assert.isNumber(numRoundedUp);
```

Sua variável `numRoundedUp` deve ser igual ao resultado de arredondar o número de ponto flutuante `3.2` para baixo até o inteiro mais próximo.

```js
assert.equal(numRoundedUp, 4);
```

Você não deve codificar o valor `4` diretamente para a variável `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Você deve registrar a variável `numRoundedUp` no console.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number between 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
