---
id: 66ece147323c91f654bf8dd9
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Nas lições anteriores, você aprendeu como obter os valores máximos e mínimos de um intervalo de números assim:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Crie uma variável chamada `maxNum` e atribua a ela o resultado de encontrar o número máximo entre os números `3, 125, 55, 24`. Então, registre o valor de `maxNum` no console.

Abaixo disso, crie uma variável chamada `minNum` e atribua a ela o resultado de encontrar o número mínimo entre os números `6, 90, 14, 90, 2`. Então, registre o valor de `minNum` no console.

# --hints--

Você deve ter uma variável chamada `maxNum`.

```js
assert.isNotNull(maxNum);
```

Sua variável `maxNum` deve ser atribuída ao resultado de encontrar o número máximo entre os números `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

Você não deve codificar o valor `125` diretamente na variável `maxNum`. Certifique-se de usar o método `Math.max()` em vez disso.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Você deve registrar o valor de `maxNum` no console.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Você deve ter uma variável chamada `minNum`.

```js
assert.isNotNull(minNum);
```

Sua variável `minNum` deve ser atribuída ao resultado de encontrar o número mínimo entre os números `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

Você não deve codificar o valor `2` diretamente na variável `minNum`. Certifique-se de usar o método `Math.min()` em vez disso.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Você deve registrar o valor de `minNum` no console.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
