---
id: 69218e46f4490b33f2481a1a
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

A fórmula para gerar um número inteiro aleatório entre dois valores é a seguinte:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Isso produzirá um resultado que é um inteiro entre dois valores.

Crie uma variável chamada `randomInt` e atribua a ela o resultado de gerar um valor entre os valores `min` e `max`.

Então, registre a variável `randomInt` para ver o resultado. Tente atualizar a página para ver resultados diferentes.

# --hints--

Você deve ter uma variável chamada `randomInt`.

```js
assert.isNotNull(randomInt);
```

Sua variável `randomInt` deve ser um número.

```js
assert.isNumber(randomInt);
```

Você deve atribuir o resultado de gerar um valor entre os valores `min` e `max`. Certifique-se de não codificar um valor fixo.  You should use the formula provided above.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Você deve registrar a variável `randomInt` no console.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
