---
id: 69218a15fe536ee1be38aefb
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

A próxima parte do workshop é revisar como gerar um número inteiro aleatório entre dois valores. 

Comece adicionando outro `console.log()` que registre a mensagem `"Now, generate a random integer between two values."` no console.

No próximo passo, você irá gerar um número inteiro aleatório entre os valores das variáveis `min` e `max` que você já criou.

# --hints--

Você deve ter um `console.log()` que registre a mensagem `"Now, generate a random integer between two values."`.

```js
assert.match(code, /console\.log\((`|'|")\s*Now,\s+generate\s+a\s+random\s+integer\s+between\s+two\s+values\.\s*\1\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

--fcc-editable-region--

--fcc-editable-region--
```
