---
id: 68e6aff5325410118e76e7be
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、ARIAを使ってアクセシブルなオーディオコントローラーを作成する練習をします。

まず、テキストが`h1`の`Audio Controls`要素を追加してください。

# --hints--

`h1`要素があるはずです。

```js
assert.exists(document.querySelector('h1'));
```

`h1`要素は`body`要素の中にあるべきです。

```js
assert.exists(document.querySelector('body > h1'));
```

`h1`要素にはテキスト`Audio Controls`があるべきです。

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
