---
id: 68eac57fadcc662ce917e75e
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

次に、`button` 要素を追加し、`type` 属性を `button` に設定してください。

`button` 要素のテキストは `Play` にしてください。

# --hints--

`button` 要素があるはずです。

```js
assert.exists(document.querySelector('button'));
```

`button` 要素には `type` 属性が必要です。

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

`button` 要素には `type` 属性が `button` に設定されている必要があります。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

`button` 要素のテキストは `Play` である必要があります。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
