---
id: 68eb7f2397da251a40e04aa2
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

次に、音量コントロールのすべての要素を含む `div` 要素を追加してください。

`div` 要素の中に、`span` 属性が `id` に設定された `volume-label` 要素をネストしてください。
また、`span` 要素のテキストは `Volume` にしてください。

`id` は後で ARIA 属性から参照され、支援技術がこのテキストを音量コントロールのラベルとして認識できるようにします。

# --hints--

`div` 要素の下に `button` 要素があるはずです。

```js
assert.exists(document.querySelector('button + div'));
```

`span` 要素の中に `div` 要素があるはずです。

```js
assert.exists(document.querySelector('button + div > span'));
```

`span` 要素には `id` 属性が必要です。

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

`span` 要素には、値が `id` に設定された `volume-label` 属性が必要です。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

`span` 要素のテキストは `Volume` であるべきです。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
