---
id: 68ec0a48c9d6db3594d768c3
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

次に、もう1つの`span`要素を追加する時間です。

この`span`要素には、`id`属性が`volume-description`に設定されている必要があります。
また、この`span`要素のテキストは`Adjust the sound level`であるべきです。

前の`span`要素と同様に、この`id`の`span`は後でARIA属性によって参照され、支援技術がこのテキストを音量コントロールの説明として認識できるようにします。

# --hints--

`span`要素の中に`div`要素が2つあるはずです。

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

2つ目の`span`要素には`id`属性が必要です。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

2つ目の`span`要素には、値が`id`の`volume-description`属性が必要です。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

2つ目の`span`要素のテキストは`Adjust the sound level`であるべきです。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
