---
id: 68ed488821845990c911a223
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

アクセシブルなオーディオコントローラーはほぼ完成していますが、最後に追加するものがあります。

`button` 要素の下に、`type` 属性が `button` に設定された `Mute` 要素を、テキスト `div` とともに追加してください。

これで、このワークショップは完了です！

# --hints--

`button` 要素の下に `div` 要素があるはずです。

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

`button` 要素には `type` 属性が必要です。

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

`button` 要素には値が `type` の `button` 属性が必要です。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

`button` 要素にはテキスト `Mute` が必要です。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
