---
id: 68ed4204ab7edc66b6cb2383
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

これまでのレッスンで覚えているように、`aria-labelledby` 属性はページ上に既存のテキストがあり、それをラベルとして使う場合に使用します。この場合、`id` 属性を`span` 要素に追加しており、ここで使われることになります。

`aria-labelledby` 要素に`input` 属性を追加し、その値を`volume-label volume-description` に設定してください。

# --hints--

`input` 要素には`aria-labelledby` 属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

`input` 要素には`aria-labelledby` 属性が必要で、その値は`volume-label volume-description` です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
