---
id: 68ed3a96e0ebcc36565b39e5
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

次に、`input`属性が`type`に設定された`range`要素を追加して、音量スライダーを作成してください。デフォルトの音量レベルを定義するために、`min`属性を`0`に、`max`属性を`100`に、`value`属性を`50`に設定してください。

# --hints--

2番目の`input`要素の下に`span`要素があるはずです。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

`input`要素には`type`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

`input`要素には値が`type`の`range`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

`input`要素には`min`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

`input`要素には値が`min`の`0`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

`input`要素には`max`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

`input`要素には値が`max`の`100`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

`input`要素には`value`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

`input`要素には値が`value`の`50`属性が必要です。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
