---
id: 694648acde178bb8202d9516
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

A saída é `False`, o que mostra que `score` não é um `int`.

Outro tipo comum de número em Python é `float`, que representa um número com decimais. Substitua `int` por `float` na chamada existente de `isinstance()` para confirmar isso.

Finalmente, imprima `score` e seu tipo de dado para completar o boletim.

# --hints--

Você deve substituir `int` por `float` na chamada existente `isinstance(score, int)`.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

Você deve imprimir `score` e `type(score)` usando um separador de vírgula.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
