---
id: 694648acde178bb8202d9512
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

O boletim também deve mostrar se o aluno está atualmente matriculado. Isso pode ser representado usando um valor booleano.

Valores booleanos representam uma condição de sim ou não e eles são frequentemente usados para tomar decisões no código. Existem apenas dois valores booleanos: `True` e `False`.

Declare uma variável chamada `is_student` e atribua a ela o valor `True`.

# --hints--

Você deve ter uma variável `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

A variável `is_student` deve armazenar o valor `True`. Não coloque o valor entre aspas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
