---
id: 69461e4f98594012939bc1c6
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste workshop, você praticará os tipos de dados do Python construindo uma impressora de boletim simples.

Como você aprendeu nas lições anteriores, as variáveis são atribuídas desta forma:

```py
greeting = 'Hello World'
```

No exemplo acima, a variável `greeting` tem o valor de uma string, que é uma sequência de caracteres cercada por aspas.

Crie uma variável chamada `name` e atribua a ela a string `Alice`. Lembre-se de cercar o valor com aspas simples ou duplas conforme mostrado no exemplo.

**Importante:** Python usa indentação (os espaços no início de uma linha) para organizar o código. Para este workshop, certifique-se de que não haja espaços extras no início de cada linha de código. Adicionar espaços extras causará um `IndentationError` e impedirá que seu código seja executado.

# --hints--

Você deve ter uma variável `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

A variável `name` deve armazenar o valor `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
