---
id: 694648acde178bb8202d9514
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora você precisa adicionar a idade do aluno ao boletim. Para isso você usará um inteiro, um dos tipos de dados numéricos em Python.

Declare uma variável chamada `age` e atribua a ela o valor inteiro `20`.

Então, imprima o valor e o tipo de dado de `age` juntos separados por vírgula. Verifique a saída no terminal que mostra o valor de `age` e seu tipo como `int` (inteiro).

# --hints--

Você deve ter uma variável `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

A variável `age` deve armazenar o valor `20`. Não coloque o valor entre aspas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Você deve imprimir `age` e `type(age)` usando um separador de vírgula.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
