---
id: 694648abde178bb8202d9511
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Você deve agora ver o nome do estudante impresso no terminal.

Python fornece uma função chamada `type()` que você pode usar para verificar o tipo de um valor.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

No exemplo acima, a saída `<class 'str'>` significa que a variável passada para a função `type()` é uma string.

Use a função `type()` com `name` como seu argumento e imprima a saída como no exemplo. Verifique a saída no terminal que mostra que `name` é do tipo `str` (string).

# --hints--

Você deve imprimir o resultado da chamada `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
