---
id: 694648acde178bb8202d9515
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Agora, adicione a pontuação do estudante.

Declare uma variável chamada `score` e atribua a ela o valor `80.5`.

Embora tanto `age` e `score` sejam números, eles podem não ser do mesmo tipo. Python fornece uma função chamada `isinstance()` para verificar isso.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Use `isinstance()` para verificar se `score` é um `int` e imprima o resultado no terminal conforme mostrado no exemplo acima.

# --hints--

Você deve ter uma variável `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

A variável `score` deve armazenar o valor `80.5`. Não coloque o valor entre aspas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Você deve imprimir o resultado de `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
