---
id: 694648acde178bb8202d9517
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

A saída é `True`, confirmando que `score` é um `float`.

Complete o boletim imprimindo o valor de `score` junto com seu tipo de dado usando uma única instrução `print()`.

# --hints--

Você deve imprimir `score` e `type(score)` usando um separador de vírgula.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
