---
id: 694648acde178bb8202d9513
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

O nome do estudante deve seguir o mesmo formato dos outros detalhes.

Remova as saídas anteriores da variável `name`. Então, imprima `name` e `type(name)` juntos em uma linha separados por vírgula como no passo anterior.

# --hints--

Você não deve ter `print(name)` no seu código.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Você não deve ter `print(type(name))` no seu código.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Você deve imprimir `name` e `type(name)` na mesma chamada de `print()` usando um separador de vírgula.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
