---
id: 6947a3e516bb8a4a681274a7
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

A função `print()` pode exibir mais de um valor ao mesmo tempo. Separe os valores com uma vírgula (`,`) para imprimi-los na mesma linha.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Imprima tanto `is_student` e `type(is_student)` na mesma linha usando uma vírgula `,` como mostrado no exemplo. Em seguida, verifique a saída no terminal que mostra o valor de `is_student` e seu tipo como `bool` (booleano).

# --hints--

Você deve imprimir `is_student` e `type(is_student)` usando um separador de vírgula.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
