---
id: 6947a0c2b6de035dd8d3eea1
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Você pode imprimir o valor de uma variável usando a função `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Você aprenderá mais sobre funções nas próximas lições. Por enquanto, saiba que uma função é um bloco reutilizável de código que pode ser chamado, ou invocado, para executar seu código e argumentos podem ser passados para ela.

No exemplo acima, `print(greeting)` é uma chamada de função e `greeting` é o argumento da função.

Consulte o exemplo e imprima a variável `name`. Verifique a saída no terminal.

# --hints--

Você deve chamar `print()` com `name` como seu argumento.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
