---
id: 66ad8d150264db3926eccfeb
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

前のレッスンでは、`let` キーワードと JavaScript で変数を宣言する方法について学びました。

ここで、`let` キーワードを使って変数を宣言する方法をおさらいしましょう。

```js
let greeting;
```

`let` キーワードを使って `bot` という名前の変数を宣言してください。

**注意**：ここでは `let` を使っていますが、ワークショップの後半で `bot` 変数の値を変更することになるからです。

# --hints--

コードには `let` キーワードが含まれているはずです。

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

`bot` という名前の変数を宣言してください。

```js
assert.isNotNull(bot);
```

変数を宣言するには `let` キーワードを使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
