---
id: 66ada3f46945763dd97f43f8
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

以前、`bot` と `botLocation` の変数を作成しました。今回はそれらを使ってコンソールに新しいメッセージを出力します。

前のレッスンでは、`+` 演算子を使って文字列を連結する方法を学びました。例えば次のように文字列をつなげます。

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

文字列と変数を連結するときは、スペースに注意する必要があります。

`botIntroduction` という変数を作成してください。

次に、`+` 演算子を使って文字列 `"My name is "` と `bot` 変数、そしてピリオド（`.`）を連結してください。

この値を `botIntroduction` 変数に代入します。

最後に、`botIntroduction` 変数をコンソールにログ出力してください。

# --hints--

`botIntroduction` という変数があるはずです。

```js
assert.isNotNull(botIntroduction);
```

`botIntroduction` 変数には文字列の値が入っているはずです。

```js
assert.isString(botIntroduction);
```

`+` 演算子を使って文字列 `"My name is "` と `bot` 変数、そしてピリオド（`.`）を連結してください。スペースに注意してください。この連結した文字列を `botIntroduction` 変数に代入します。

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

コードには4つ目の `console.log()` 文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

`console` 文は `botIntroduction` 変数を出力するはずです。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
