---
id: 66adbc8a9793d64250f7e609
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

前のレッスンでは、変数に値を再代入する方法を学びました。

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

再代入を使って、文字列 `"professorBot"` を `bot` 変数に代入してください。

# --hints--

文字列 `"professorBot"` を `bot` 変数に代入する必要があります。

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
