---
id: 66adc82170acb1464bd348eb
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

ボットからの最後の数メッセージは、ボットのお気に入りの話題にフォーカスします。

まず、`favoriteSubject` という新しい変数を作成し、文字列 `"Computer Science"` を代入してください。

# --hints--

`favoriteSubject` という変数があるはずです。

```js
assert.isNotNull(favoriteSubject);
```

`favoriteSubject` 変数は文字列であるべきです。

```js
assert.isString(favoriteSubject);
```

`"Computer Science"` 変数に文字列 `favoriteSubject` を代入してください。

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
