---
id: 66adcf383276814776aba3ca
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

次に、`favoriteSubjectSentence` という変数を作成してください。

文字列の連結を使って、文字列 `"My favorite subject is "` と `favoriteSubject` 変数をピリオドでつなげてください。

その結果を `favoriteSubjectSentence` 変数に代入してください。

そして、`favoriteSubjectSentence` の値をコンソールに出力してください。

# --hints--

`favoriteSubjectSentence` という変数があるはずです。

```js
assert.isNotNull(favoriteSubjectSentence);
```

`favoriteSubjectSentence` 変数は文字列であるべきです。

```js
assert.isString(favoriteSubjectSentence);
```

文字列 `"My favorite subject is "` と `favoriteSubject` 変数をピリオドでつなげるために文字列の連結を使うべきです。スペースに注意してください。

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

コード内に `console` 文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

`favoriteSubjectSentence` の値をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
