---
id: 66ad9c633fa26d3c1475eae3
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

次に、ボットからの別のメッセージを追加しましょう。

コードにもう1つ`console`文を追加して、メッセージ`"Allow me to introduce myself."`をログに記録してください。

# --hints--

コードには3つ目の`console.log()`文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

`console`文にはメッセージ`"Allow me to introduce myself."`を含めてください。引用符を忘れないでください。

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
