---
id: 66adc42868cab843ccee87d9
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

新しい `bot` の値を確認する時間です。

まず、`nicknameIntroduction` という新しい変数を作成してください。

文字列の連結を使って、文字列 `"My nickname is "` と `bot` 変数をピリオド（`.`）でつなげてください。

その結果の文字列を `nicknameIntroduction` 変数に代入してください。

そして、`nicknameIntroduction` の値をコンソールに出力してください。

# --hints--

`nicknameIntroduction` という変数があるはずです。

```js
assert.isNotNull(nicknameIntroduction);
```

`nicknameIntroduction` 変数は文字列であるべきです。

```js
assert.isString(nicknameIntroduction);
```

文字列 `"My nickname is "` と `bot` 変数をピリオド（`.`）で連結してください。スペースに注意してください。

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

コード内に `console` 文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

`nicknameIntroduction` の値をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
