---
id: 66adc5919c1853448f119d11
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

今、ボットはニックネームを変更したがっています。

再代入を使って、文字列 `"awesomeTeacherBot"` を `bot` 変数に代入してください。

# --hints--

文字列 `"awesomeTeacherBot"` を `bot` 変数に代入する必要があります。

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
