---
id: 66adb844118ba74107ce771f
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

次のボットからのメッセージは、ボットの場所に関するものです。

`botLocationSentence` という変数を作成してください。

次に、`+` 演算子を使って文字列の連結を行い、文字列 `"I live in "` と `botLocation` 変数、その後にピリオド（`.`）をつなげてください。

この値を `botLocationSentence` 変数に代入してください。

最後に、`botLocationSentence` の値をコンソールに出力してください。

# --hints--

`botLocationSentence` という変数があるはずです。

```js
assert.isNotNull(botLocationSentence);
```

`botLocationSentence` 変数には文字列の値が入っているはずです。

```js
assert.isString(botLocationSentence);
```

`+` 演算子を使って文字列 `"I live in "` と `botLocation` 変数、その後にピリオド（`.`）を連結してください。スペースに注意してください。

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

コード内に `console` 文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

`console` 文は `botLocationSentence` 変数を出力するはずです。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
