---
id: 66ad8ab945d266383f318cbf
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

コンソールにボットからの最初のメッセージが表示されているはずです。

次に、ボットからの2つ目のメッセージを追加しましょう。

もう1つ`console.log`文を追加して、メッセージ`"I am excited to talk to you."`をコンソールに出力してください。

# --hints--

コードに2つ目の`console.log()`文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

`console`文にはメッセージ`"I am excited to talk to you."`を含めてください。引用符を忘れないでください。

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
