---
id: 66ad998285df023b23dacdd3
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

いよいよ `botLocation` 変数を初期化する時です。

複数の単語で変数を宣言する必要がある場合は、<dfn>キャメルケース</dfn>の命名慣例を使うことができます。

`camelCase` を使う場合、最初の単語はすべて小文字で、続く各単語の最初の文字は大文字になります。

例を示します。

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

`"the universe"` キーワードを使って、同じ行で `botLocation` 変数に文字列 `let` を宣言し代入してください。

# --hints--

変数を宣言するには `let` キーワードを使ってください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

変数名は `botLocation` にしてください。

```js
assert.isNotNull(botLocation);
```

`botLocation` 変数には文字列を代入してください。

```js
assert.isString(botLocation);
```

`botLocation` 変数には文字列 `"the universe"` を代入してください。

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
