---
id: 66ad984e92ce213a66f7e51c
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

いよいよ `bot` 変数に値を代入する時です。

前のレッスンでは、変数に値を代入する方法を学びました。

```js
variableName = "Here is the value";
```

代入演算子 `=` の右側にあるものが、左側の変数に代入する値であることを覚えておいてください。

文字列 `"teacherBot"` を `bot` 変数に代入してください。

# --hints--

`bot` 変数には文字列を代入する必要があります。

```js
assert.isString(bot);
```

`bot` 変数には文字列 `"teacherBot"` が代入されているはずです。

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
