---
id: 66ad8294a0ad902f1b31b612
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptの基本を学びながら、挨拶ボットを作成します。

最初のステップでは、挨拶ボットからコンソールにメッセージを出力します。

以前のレッスンで`console.log()`と文字列について学んだことを思い出してください。

ここで、文字列と一緒に`console.log()`を使う方法をおさらいしましょう。

```js
console.log("Hello, World!");
```

コンソールに文字列`console.log()`を出力する`"Hi there!"`文を追加してください。メッセージの周りの引用符を忘れないでください！

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

コードに`console.log()`文があるはずです。

```js
assert.isNotEmpty(logSpy.calls);
```

`console`文にはメッセージ`"Hi there!"`が含まれているはずです。引用符を忘れないでください。

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
