---
id: 66adc6046acc18453decf577
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

ボットの新しいニックネームを見るには、新しいメッセージをコンソールにログ出力する必要があります。

`newNicknameGreeting` という名前の新しい変数を作成してください。

次に、文字列連結を使って、文字列 `"I love my nickname but I wish people would call me "` と `bot` 変数をピリオドでつなげてください。

その結果を `newNicknameGreeting` 変数に代入してください。

最後に、`newNicknameGreeting` の値をコンソールにログ出力してください。

# --hints--

`newNicknameGreeting` という変数があるはずです。

```js
assert.isNotNull(newNicknameGreeting);
```

`newNicknameGreeting` 変数は文字列であるべきです。

```js
assert.isString(newNicknameGreeting);
```

文字列連結を使って、文字列 `"I love my nickname but I wish people would call me "` と `bot` 変数をピリオドでつなげてください。

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

コード内に `console` 文があるはずです。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

`newNicknameGreeting` の値をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
