---
id: 60eebd07ea685b0e777b5583
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

등록 양식 프로젝트에 오신 것을 환영합니다! 문서 상단에 `!DOCTYPE html` 선언을 추가하여 브라우저가 어떤 유형의 문서를 읽고 있는지 알 수 있게 하세요.

# --hints--

코드에 `DOCTYPE` 참조가 포함되어야 합니다.

```js
assert(code.match(/<!DOCTYPE/gi));
```

`DOCTYPE` 참조 뒤에는 공백을 포함해야 합니다.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

문서 유형을 `html`로 정의해야 합니다.

```js
assert(code.match(/html/gi));
```

`DOCTYPE` 선언은 유형 뒤에 `>`로 닫아야 합니다.

```js
assert(code.match(/html\s*>/gi));
```

`DOCTYPE` 선언은 HTML의 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
