---
id: 60f0286404aefb0562a4fdf9
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`title` 요소 안에 `meta`와 `head` 요소를 추가하세요. 프로젝트 제목을 `Registration Form`로 지정하고, `charset` 요소에 `utf-8` 속성을 값 `meta`로 추가하세요.

# --hints--

코드에 `title` 요소가 있어야 합니다.

```js
const title = document.querySelector('title');
assert.exists(title);
```

`title` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('head > title'));
```

프로젝트 제목은 `Registration Form`이어야 합니다. 제목의 대소문자와 철자를 정확히 지키세요.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

`meta` 요소 안에 `head` 요소를 만들어야 합니다.

```js
assert.exists(document.querySelector('head > meta'));
```

`meta` 요소에 `charset` 속성을 값 `utf-8`로 지정해야 합니다.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

`meta` 요소는 종료 태그가 없는 빈 요소여야 합니다 `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
