---
id: 60f027099a15b00485563dd2
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

`DOCTYPE` 아래에 `html` 속성이 `lang`로 설정된 `en` 요소를 추가해서 코드를 작성할 공간을 만드세요.

# --hints--

`DOCTYPE` 선언은 HTML 맨 앞에 있어야 합니다.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

`html` 요소는 여는 태그가 있어야 하며, `lang` 속성의 값은 `en`이어야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 요소는 닫는 태그가 있어야 합니다.

```js
assert(code.match(/<\/html\s*>/));
```

`html` 태그는 올바른 순서여야 합니다.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

`html` 요소는 하나만 있어야 합니다.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
