---
id: 60f02e7361b68405e27b62a5
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`body` 안에 `h1` 텍스트가 들어간 `Registration Form`을 추가하여 콘텐츠에 제목 문맥을 제공하세요.

# --hints--

`h1`을 `body` 안에 추가해야 합니다.

```js
assert.exists(document.querySelector('body > h1'));
```

`h1`에 `Registration Form` 텍스트를 넣어야 합니다.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
