---
id: 60f80e0081e0f2052ae5b505
title: 19단계
challengeType: 0
dashedName: step-19
---

# --description--

`type` 요소의 `input` 속성을 지정하는 것은 브라우저가 어떤 종류의 데이터를 예상해야 하는지 알기 위해 중요합니다. `type`가 지정되지 않으면 브라우저는 기본값으로 `text`를 사용합니다.

처음 두 `input` 요소에는 `type` 값을 가진 `text` 속성을, 세 번째에는 `type` 값을 가진 `email` 속성을, 네 번째에는 `type` 값을 가진 `password` 속성을 지정하세요.

`email` 유형은 도메인에 `@`와 `.`가 포함된 이메일만 허용합니다.
`password` 유형은 입력을 가리고, 사이트가 HTTPS를 사용하지 않으면 경고합니다.

# --hints--

첫 번째 `input` 요소에는 `type` 값을 가진 `text` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

두 번째 `input` 요소에는 `type` 값을 가진 `text` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

세 번째 `input` 요소에는 `type` 값을 가진 `email` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

네 번째 `input` 요소에는 `type` 값을 가진 `password` 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
        <label for="email">Enter Your Email: <input id="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
