---
id: 60f852f645b5310a8264f555
title: 24단계
challengeType: 0
dashedName: step-24
---

# --description--

등록 양식의 다음 부분으로 넘어가겠습니다. 이 섹션에서는 사용자가 개설하는 계좌 유형을 묻습니다.

두 번째 `label`에 `fieldset` 요소 두 개를 추가하는 것부터 시작하세요.

# --hints--

두 번째 `label`에 `fieldset` 요소 두 개를 추가해야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset')?.[1]?.querySelectorAll('label')?.length, 2);
```

`label` 요소들은 형제 요소여야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(2)')?.querySelector('label + label'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
