---
id: 60f83e7bfc09900959f41e20
title: 22단계
challengeType: 0
dashedName: step-22
---

# --description--

일부 `type` 속성 값에는 내장된 폼 유효성 검사가 포함되어 있습니다. 예를 들어, `type="email"`은 값이 유효한 이메일 주소여야 한다는 조건이 있습니다.

`input` 요소에 `minlength` 속성을 값 `8`로 추가하여 비밀번호에 대한 사용자 지정 유효성 검사를 추가하세요. 이렇게 하면 8자 미만의 입력이 제출되는 것을 방지합니다.

# --hints--

비밀번호 `input` 요소에 `minlength` 속성을 추가하세요.

```js
assert.notEqual(document.querySelector('input[type="password"]')?.minLength, -1);
```

`minlength` 속성에 `8` 값을 지정하세요.

```js
assert.equal(document.querySelector('input[type="password"]')?.minLength, 8);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
