---
id: 60f8618d191b940d62038513
title: 27단계
challengeType: 0
dashedName: step-27
---

# --description--

한 번에 하나의 라디오 입력만 선택되게 하세요. 하지만 폼은 라디오 입력들이 서로 관련되어 있다는 것을 알지 못합니다.

라디오 입력들을 관련시키려면, 같은 `name` 속성을 값 `account-type`으로 지정하세요. 이제 두 라디오 입력을 동시에 선택할 수 없습니다.

# --hints--

첫 번째 라디오 입력에 `name` 속성을 값 `account-type`으로 지정하세요.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.name, 'account-type');
```

두 번째 라디오 입력에 `name` 속성을 값 `account-type`으로 지정하세요.

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.name, 'account-type');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label><input type="radio" /> Personal</label>
        <label><input type="radio" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
