---
id: 60f81616cff80508badf9ad5
title: 21단계
challengeType: 0
dashedName: step-21
---

# --description--

이 시점에서 폼을 제출할 수 있어야 합니다. 하지만, 별다른 변화가 없다는 것을 알 수 있습니다.

폼을 더 상호작용적으로 만들려면, 첫 번째 `required` 안에 있는 `input` 요소들에 `fieldset` 속성을 추가하세요.

이제 필수 필드를 채우지 않고 폼을 제출하려고 하면, 오류 메시지가 표시됩니다.

# --hints--

첫 번째 `input` 요소에 `required` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelector('input')?.required, true);
```

두 번째 `input` 요소에 `required` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

세 번째 `input` 요소에 `required` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

네 번째 `input` 요소에 `required` 속성을 추가해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

`submit` `input`에는 `required` 속성을 추가하지 않아야 합니다.

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
