---
id: 60fad1cafcde010995e15306
title: 42단계
challengeType: 0
dashedName: step-42
---

# --description--

폼 제출 시, 각 제출 가능한 요소에 `name` 속성을 부여하는 것이 유용하며 좋은 습관입니다. 이 속성은 폼 제출에서 요소를 식별하는 데 사용됩니다.

이미 이름을 지정한 두 개의 `radio` 입력을 제외하고, 각 제출 가능한 요소에 고유한 `name` 속성을 원하는 대로 부여하세요.

# --hints--

첫 이름을 입력받는 `input`에 `name` 속성을 부여해야 합니다. _참고로 저는 `first-name`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

성(last name)을 입력받는 `input`에 `name` 속성을 부여해야 합니다. _참고로 저는 `last-name`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

`email`에 `name` 속성을 부여해야 합니다. _참고로 저는 `email`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

`password`에 `name` 속성을 부여해야 합니다. _참고로 저는 `password`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

`checkbox`에 `name` 속성을 부여해야 합니다. _참고로 저는 `terms`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

`file`에 `name` 속성을 부여해야 합니다. _참고로 저는 `file`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

`number`에 `name` 속성을 부여해야 합니다. _참고로 저는 `age`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

`select` 요소에 `name` 속성을 부여해야 합니다. _참고로 저는 `referrer`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

`textarea` 요소에 `name` 속성을 부여해야 합니다. _참고로 저는 `bio`_를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

어떤 `option` 요소에도 `name` 속성을 부여하지 마세요.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

어떤 `label` 요소에도 `name` 속성을 부여하지 마세요.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

어떤 `fieldset` 요소에도 `name` 속성을 부여하지 마세요.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      --fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
