---
id: 60fac4095512d3066053d73c
title: 34단계
challengeType: 0
dashedName: step-34
---

# --description--

`select` 요소로 폼에 드롭다운을 추가하는 것은 쉽습니다. `select` 요소는 `option` 요소 그룹을 담는 컨테이너 역할을 하며, `option` 요소는 각 드롭다운 옵션의 레이블 역할을 합니다. 두 요소 모두 닫는 태그가 필요합니다.

마지막 `select` 요소 아래에 `label` 요소를 추가하는 것부터 시작하세요. 그런 다음 `option` 요소 안에 5개의 `select` 요소를 중첩하세요.

# --hints--

세 번째 `select`에 `fieldset` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > select'));
```

`option` 요소 안에 5개의 `select` 요소를 중첩해야 합니다.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(3) > select > option')?.length, 5);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset>
        <label>Upload a profile picture: <input type="file" /></label>
        <label>Input your age (years): <input type="number" min="13" max="120" /></label>

      </fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
