---
id: 60fab8367d35de04e5cb7929
title: 31단계
challengeType: 0
dashedName: step-31
---

# --description--

새로 추가한 `I accept the terms and conditions` 안에서 `input` 요소 바로 뒤에 `label` 텍스트를 추가하세요. 그런 다음 `terms and conditions` 텍스트를 다음 위치에 링크하세요:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

세 번째 `I accept the terms and conditions` 다음에 있는 `label`에 `fieldset` 텍스트를 추가해야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

약관에 링크하려면 `a` 요소를 사용해야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

새 텍스트를 `input` 안의 `label` 요소 바로 뒤에 넣어야 합니다.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

`a` 요소에 `href`을 `https://www.freecodecamp.org/news/terms-of-service/`로 지정해야 합니다.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

`a` 요소는 `terms and conditions` 텍스트만 감싸야 합니다.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

앵커 요소 안의 텍스트에 불필요한 앞뒤 공백이 있습니다. 앵커 텍스트 내 공백은 `terms`, `and`, `conditions` 단어 사이에만 있어야 합니다.

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
