---
id: 62cc5b1779e4d313466f73c5
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`link` 요소 안에 `head` 요소를 중첩하세요. `rel` 속성을 값 `stylesheet`으로, `href` 속성을 값 `styles.css`로 지정하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector('link'));
```

기존 `head` 태그를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

`link` 요소가 하나 있어야 합니다.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 요소는 빈 요소여야 하며, 닫는 태그 `</link>`가 없어야 합니다.

```js
assert.notMatch(code, /<\/link>/);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 요소는 값 `rel`인 `stylesheet` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소는 값 `href`인 `styles.css` 속성을 가져야 합니다.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
