---
id: 5a24c314108439a4d403616a
title: Ein Array als Props übergeben
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

Die letzte Aufgabe zeigte, wie man Informationen von einer übergeordneten Komponente an eine untergeordnete Komponente als `props` oder Eigenschaften übergibt. Diese Aufgabe betrachtet, wie Arrays als `props` übergeben werden können. Um ein Array an ein JSX-Element zu übergeben, muss es als JavaScript behandelt und in geschweifte Klammern eingeschlossen werden.

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

Die untergeordnete Komponente hat dann Zugriff auf die Array-Eigenschaft `colors`. Array-Methoden wie `join()` können beim Zugriff auf die Eigenschaft verwendet werden.

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
```

Dies verbindet alle `colors` Array-Elemente zu einem durch Kommas getrennten String und erzeugt: `<p>green, blue, red</p>`. Später lernen wir weitere gängige Methoden kennen, um Arrays von Daten in React zu rendern.

# --instructions--

Im Code-Editor befinden sich `List` und `ToDo` Komponenten. Wenn Sie jede `List` aus der `ToDo` Komponente rendern, übergeben Sie eine `tasks` Eigenschaft, der ein Array von To-do-Aufgaben zugewiesen ist, zum Beispiel `["walk dog", "workout"]`. Greifen Sie dann in der `tasks` Komponente auf dieses `List` Array zu und zeigen Sie dessen Wert innerhalb des `p` Elements an. Verwenden Sie `join(", ")`, um das `props.tasks` Array im `p` Element als durch Kommas getrennte Liste anzuzeigen. Die heutige Liste sollte mindestens 2 Aufgaben enthalten, die von morgen mindestens 3 Aufgaben.

# --hints--

Die `ToDo` Komponente sollte ein einzelnes äußeres `div` zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

Das dritte Kind der `ToDo` Komponente sollte eine Instanz der `List` Komponente sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

Das fünfte Kind der `ToDo` Komponente sollte eine Instanz der `List` Komponente sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

Beide Instanzen der `List` Komponente sollten eine Eigenschaft namens `tasks` haben, und `tasks` sollte vom Typ Array sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

Die erste `List` Komponente, die die Aufgaben für heute darstellt, sollte 2 oder mehr Elemente enthalten.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

Die zweite `List` Komponente, die die Aufgaben für morgen darstellt, sollte 3 oder mehr Elemente enthalten.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

Die `List` Komponente sollte den Wert aus der `tasks` Prop im `p` Tag rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
