---
id: 5a24c314108439a4d4036181
title: Einführung in inline styles
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

Es gibt weitere komplexe Konzepte, die Ihrem React-Code mächtige Fähigkeiten hinzufügen. Aber vielleicht fragen Sie sich, wie Sie die JSX-Elemente, die Sie in React erstellen, auf einfachere Weise stylen können. Wahrscheinlich wissen Sie, dass es nicht genau so funktioniert wie bei HTML, wegen <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">der Art und Weise, wie Sie Klassen auf JSX-Elemente anwenden</a>.

Wenn Sie Styles aus einem Stylesheet importieren, ist es kaum anders. Sie weisen Ihrem JSX-Element eine Klasse mit dem `className`-Attribut zu und wenden die Styles in Ihrem Stylesheet auf diese Klasse an. Eine andere Möglichkeit ist, inline styles anzuwenden, die in der ReactJS-Entwicklung sehr verbreitet sind.

Sie wenden inline styles auf JSX-Elemente ähnlich an wie in HTML, aber mit einigen JSX-Unterschieden. Hier ist ein Beispiel für einen inline style in HTML:

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

JSX-Elemente verwenden das `style`-Attribut, aber wegen der Art, wie JSX transpiliert wird, können Sie den Wert nicht auf eine `string` setzen. Stattdessen setzen Sie ihn gleich einem JavaScript-`object`. Hier ein Beispiel:

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

Beachten Sie, wie wir die `fontSize`-Eigenschaft in CamelCase schreiben? Das liegt daran, dass React keine kebab-case-Schlüssel im Style-Objekt akzeptiert. React wendet für uns den korrekten Eigenschaftsnamen im HTML an.

# --instructions--

Fügen Sie dem `style` im Code-Editor ein `div`-Attribut hinzu, um dem Text die Farbe Rot und eine Schriftgröße von `72px` zu geben.

Beachten Sie, dass Sie die Schriftgröße optional als Zahl ohne die Einheit `px` angeben oder als `"72px"` schreiben können.
# --hints--

Die Komponente sollte ein `div`-Element rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Das `div`-Element sollte die Farbe `red` haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

Das `div`-Element sollte eine Schriftgröße von `72px` haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
