---
id: 5a24c314108439a4d4036166
title: React-Komponenten zusammensetzen
challengeType: 6
forumTopicId: 301381
dashedName: compose-react-components
---

# --description--

Da die Aufgaben zunehmend komplexere Zusammensetzungen mit React-Komponenten und JSX verwenden, gibt es einen wichtigen Punkt zu beachten. Das Rendern von ES6-Klassenkomponenten innerhalb anderer Komponenten unterscheidet sich nicht vom Rendern der einfachen Komponenten, die Sie in den letzten Aufgaben verwendet haben. Sie können JSX-Elemente, zustandslose funktionale Komponenten und ES6-Klassenkomponenten innerhalb anderer Komponenten rendern.

# --instructions--

Im Code-Editor rendert die `TypesOfFood`-Komponente bereits eine Komponente namens `Vegetables`. Außerdem gibt es die `Fruits`-Komponente aus der letzten Aufgabe.

Verschachteln Sie zwei Komponenten innerhalb von `Fruits` – zuerst `NonCitrus` und dann `Citrus`. Beide Komponenten werden Ihnen im Hintergrund bereitgestellt. Verschachteln Sie anschließend die `Fruits`-Klassenkomponente in die `TypesOfFood`-Komponente, unterhalb des `h1`-Überschrift-Elements und oberhalb von `Vegetables`. Das Ergebnis sollte eine Reihe verschachtelter Komponenten sein, die zwei verschiedene Komponententypen verwendet.

# --before-all--

```js
globalThis.NonCitrus = function NonCitrus() {
  return React.createElement('div', null);
};
globalThis.Citrus = function Citrus() {
  return React.createElement('div', null);
};
globalThis.Vegetables = function Vegetables() {
  return React.createElement('div', null);
};
```

# --hints--

Die `TypesOfFood`-Komponente sollte ein einzelnes `div`-Element zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Die `TypesOfFood`-Komponente sollte die `Fruits`-Komponente zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Die `Fruits`-Komponente sollte die `NonCitrus`-Komponente und die `Citrus`-Komponente zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return (
      mockedComponent.find('Fruits').children().find('NonCitrus').length ===
        1 &&
      mockedComponent.find('Fruits').children().find('Citrus').length === 1
    );
  })()
);
```

Die `TypesOfFood`-Komponente sollte die `Vegetables`-Komponente unterhalb der `Fruits`-Komponente zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
        <Vegetables />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Fruits extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h2>Fruits:</h2>
        { /* Change code below this line */ }
        <NonCitrus />
        <Citrus />
        { /* Change code above this line */ }
      </div>
    )
  }
}

class TypesOfFood extends React.Component {
  constructor(props) {
     super(props);
  }
    render() {
      return (
        <div>
        <h1>Types of Food:</h1>
          { /* Change code below this line */ }
          <Fruits />
          { /* Change code above this line */ }
          <Vegetables />
        </div>
      );
    }
};
```
