---
id: 5a24c314108439a4d403616e
title: Props mit this.props zugreifen
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

Die letzten Aufgaben behandelten die grundlegenden Methoden, um Props an Kind-Komponenten zu übergeben. Aber was ist, wenn die Kind-Komponente, der Sie eine Prop übergeben, eine ES6-Klassenkomponente ist und keine zustandslose funktionale Komponente? Die ES6-Klassenkomponente verwendet eine etwas andere Konvention, um auf Props zuzugreifen.

Immer wenn Sie innerhalb einer Klassenkomponente auf sich selbst verweisen, verwenden Sie das Schlüsselwort `this`. Um innerhalb einer Klassenkomponente auf Props zuzugreifen, setzen Sie vor den Code, mit dem Sie darauf zugreifen, `this`. Wenn eine ES6-Klassenkomponente zum Beispiel eine Prop namens `data` hat, schreiben Sie in JSX `{this.props.data}`.

# --instructions--

Rendern Sie eine Instanz der `Welcome`-Komponente in der übergeordneten Komponente `App`. Geben Sie hier `Welcome` eine Prop namens `name` und weisen Sie ihr einen String-Wert zu. Greifen Sie innerhalb des Kindes `Welcome` auf die Prop `name` innerhalb der `strong`-Tags zu.

# --hints--

Die `App`-Komponente sollte ein einzelnes `div`-Element zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Das Kind von `App` sollte die `Welcome`-Komponente sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

Die `Welcome`-Komponente sollte eine Prop namens `name` haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

Die `Welcome`-Komponente sollte den String anzeigen, den Sie als Prop `name` übergeben, und zwar innerhalb der `strong`-Tags.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
