---
id: 5a24c314108439a4d4036179
title: Ein kontrolliertes Formular erstellen
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

Die letzte Aufgabe zeigte, dass React den internen Zustand für bestimmte Elemente wie `input` und `textarea` steuern kann, wodurch sie zu kontrollierten Komponenten werden. Dies gilt auch für andere Formularelemente, einschließlich des regulären HTML-`form`-Elements.

# --instructions--

Die `MyForm`-Komponente ist mit einem leeren `form` und einem Submit-Handler eingerichtet. Der Submit-Handler wird aufgerufen, wenn das Formular abgesendet wird.

Wir haben einen Button hinzugefügt, der das Formular absendet. Sie sehen, dass er das `type` auf `submit` gesetzt hat, was anzeigt, dass er der Button ist, der das Formular steuert. Fügen Sie das `input`-Element in das `form` ein und setzen Sie seine `value`- und `onChange()`-Attribute wie in der letzten Aufgabe. Dann sollten Sie die `handleSubmit`-Methode vervollständigen, sodass sie die Komponentenzustands-Eigenschaft `submit` auf den aktuellen Eingabewert im lokalen `state` setzt.

**Hinweis:** Sie müssen auch `event.preventDefault()` im Submit-Handler aufrufen, um das Standardverhalten des Formularabsendens zu verhindern, das die Webseite aktualisieren würde. Zum Komfort der Lernenden wurde das Standardverhalten hier deaktiviert, um zu verhindern, dass Aktualisierungen den Code der Aufgabe zurücksetzen.

Erstellen Sie schließlich ein `h1`-Tag nach dem `form`, das den `submit`-Wert aus dem `state` der Komponente rendert. Sie können dann im Formular tippen und den Button klicken (oder Enter drücken), und Sie sollten Ihre Eingabe auf der Seite gerendert sehen.

# --hints--

`MyForm` sollte ein `div`-Element zurückgeben, das ein `form` und ein `h1`-Tag enthält. Das Formular sollte ein `input` und ein `button` enthalten.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

Der Zustand von `MyForm` sollte mit den Eigenschaften `input` und `submit` initialisiert werden, die beide auf leere Strings gesetzt sind.

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

Das Tippen im `input`-Element sollte die `input`-Eigenschaft des Komponentenzustands aktualisieren.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

Das Absenden des Formulars sollte `handleSubmit` ausführen, das die `submit`-Eigenschaft im Zustand auf die aktuelle Eingabe setzt.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit` sollte `event.preventDefault` aufrufen.

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

Das `h1`-Überschriften-Element sollte den Wert des `submit`-Felds aus dem Zustand der Komponente rendern.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
