---
id: 5a24c314108439a4d4036187
title: Verwenden Sie einen ternären Ausdruck für bedingtes Rendern
challengeType: 6
forumTopicId: 301414
dashedName: use-a-ternary-expression-for-conditional-rendering
---

# --description--

Bevor wir zu dynamischen Render-Techniken übergehen, gibt es noch eine letzte Möglichkeit, eingebaute JavaScript-Bedingungen zu verwenden, um das Gewünschte zu rendern: den <dfn>ternären Operator</dfn>. Der ternäre Operator wird oft als Abkürzung für `if/else` Anweisungen in JavaScript genutzt. Sie sind nicht ganz so umfangreich wie traditionelle `if/else` Anweisungen, aber bei React-Entwicklern sehr beliebt. Ein Grund dafür ist, dass aufgrund der Art, wie JSX kompiliert wird, `if/else` Anweisungen nicht direkt in JSX-Code eingefügt werden können. Das ist Ihnen vielleicht schon bei einigen vorherigen Aufgaben aufgefallen – wenn eine `if/else` Anweisung erforderlich war, stand sie immer *außerhalb* der `return` Anweisung. Ternäre Ausdrücke können eine ausgezeichnete Alternative sein, wenn Sie bedingte Logik innerhalb Ihres JSX implementieren möchten. Denken Sie daran, dass ein ternärer Operator aus drei Teilen besteht, Sie aber mehrere ternäre Ausdrücke miteinander kombinieren können. Hier ist die grundlegende Syntax:

```jsx
condition ? expressionIfTrue : expressionIfFalse;
```

# --instructions--

Im Code-Editor sind drei Konstanten innerhalb der `CheckUserAge` Komponente in der `render()` Methode definiert. Sie heißen `buttonOne`, `buttonTwo` und `buttonThree`. Jede von ihnen ist einem einfachen JSX-Ausdruck zugewiesen, der ein Button-Element darstellt. Initialisieren Sie zuerst den Zustand von `CheckUserAge` mit `input` und `userAge`, die beide auf Werte eines leeren Strings gesetzt sind.

Sobald die Komponente Informationen auf der Seite rendert, sollten Benutzer eine Möglichkeit haben, mit ihr zu interagieren. Richten Sie innerhalb der `return` Anweisung der Komponente einen ternären Ausdruck ein, der folgende Logik umsetzt: Wenn die Seite zum ersten Mal geladen wird, rendern Sie die Absenden-Schaltfläche `buttonOne` auf der Seite. Wenn dann ein Benutzer sein Alter eingibt und die Schaltfläche klickt, rendern Sie eine andere Schaltfläche basierend auf dem Alter. Wenn ein Benutzer eine Zahl kleiner als `18` eingibt, rendern Sie `buttonThree`. Wenn ein Benutzer eine Zahl größer oder gleich `18` eingibt, rendern Sie `buttonTwo`.

# --hints--

Die `CheckUserAge` Komponente sollte mit einem einzigen `input` Element und einem einzigen `button` Element rendern.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('input')
    .length === 1 &&
    Enzyme.mount(React.createElement(CheckUserAge)).find('div').find('button')
      .length === 1
);
```

Der Zustand der `CheckUserAge` Komponente sollte mit einer Eigenschaft `userAge` und einer Eigenschaft `input` initialisiert werden, die beide auf einen leeren String gesetzt sind.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).state().input === '' &&
    Enzyme.mount(React.createElement(CheckUserAge)).state().userAge === ''
);
```

Wenn die `CheckUserAge` Komponente zum ersten Mal im DOM gerendert wird, sollte der innere Text von `button` „Submit“ lauten.

```js
assert(
  Enzyme.mount(React.createElement(CheckUserAge)).find('button').text() ===
    'Submit'
);
```

Wenn eine Zahl kleiner als 18 in das `input` Element eingegeben wird und die `button` geklickt wird, sollte der innere Text von `button` `You Shall Not Pass` lauten.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter3AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '3' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter17AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '17' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge3 = enter3AndClickButton();
  const userAge17 = enter17AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge3 === 'You Shall Not Pass' &&
      userAge17 === 'You Shall Not Pass'
  );
})();
```

Wenn eine Zahl größer oder gleich 18 in das `input` Element eingegeben wird und die `button` geklickt wird, sollte der innere Text von `button` `You May Enter` lauten.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const initialButton = mockedComponent.find('button').text();
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter35AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '35' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const userAge35 = enter35AndClickButton();
  assert(
    initialButton === 'Submit' &&
      userAge18 === 'You May Enter' &&
      userAge35 === 'You May Enter'
  );
})();
```

Sobald eine Zahl abgesendet wurde und der Wert von `input` erneut geändert wird, sollte `button` wieder `Submit` anzeigen.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(CheckUserAge));
  const enter18AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '18' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const changeInputDontClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '5' } });
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const enter10AndClickButton = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: '10' } });
    mockedComponent.find('button').simulate('click');
    mockedComponent.update();
    return mockedComponent.find('button').text();
  };
  const userAge18 = enter18AndClickButton();
  const changeInput1 = changeInputDontClickButton();
  const userAge10 = enter10AndClickButton();
  const changeInput2 = changeInputDontClickButton();
  assert(
    userAge18 === 'You May Enter' &&
      changeInput1 === 'Submit' &&
      userAge10 === 'You Shall Not Pass' &&
      changeInput2 === 'Submit'
  );
})();
```

Ihr Code sollte keine `if/else` Anweisungen enthalten.

```js
assert(
  new RegExp(/(\s|;)if(\s|\()/).test(
    code
  ) === false
);
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class CheckUserAge extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userAge: '',
      input: ''
    };
    this.submit = this.submit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(e) {
    this.setState({
      input: e.target.value,
      userAge: ''
    });
  }
  submit() {
    this.setState(state => ({
      userAge: state.input
    }));
  }
  render() {
    const buttonOne = <button onClick={this.submit}>Submit</button>;
    const buttonTwo = <button>You May Enter</button>;
    const buttonThree = <button>You Shall Not Pass</button>;
    return (
      <div>
        <h3>Enter Your Age to Continue</h3>
        <input
          style={inputStyle}
          type='number'
          value={this.state.input}
          onChange={this.handleChange}
        />
        <br />
        {this.state.userAge === ''
          ? buttonOne
          : this.state.userAge >= 18
          ? buttonTwo
          : buttonThree}
      </div>
    );
  }
}
```
