---
id: 5a24c314108439a4d403617e
title: Ereignis-Listener hinzufügen
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

Die `componentDidMount()`-Methode ist auch der beste Ort, um alle Ereignis-Listener anzuhängen, die Sie für bestimmte Funktionen hinzufügen müssen. React stellt ein synthetisches Ereignissystem bereit, das das native Ereignissystem der Browser umschließt. Das bedeutet, dass sich das synthetische Ereignissystem unabhängig vom Browser des Nutzers genau gleich verhält – selbst wenn sich die nativen Ereignisse zwischen verschiedenen Browsern unterschiedlich verhalten können.

Sie haben bereits einige dieser synthetischen Ereignis-Handler wie `onClick()` verwendet. Das synthetische Ereignissystem von React ist ideal für die meisten Interaktionen, die Sie mit DOM-Elementen verwalten. Wenn Sie jedoch einen Ereignis-Handler an die Objekte document oder window anhängen möchten, müssen Sie dies direkt tun.

# --instructions--

Hängen Sie in der `componentDidMount()`-Methode einen Ereignis-Listener für `keydown`-Ereignisse an und lassen Sie diese Ereignisse den Callback `handleKeyPress()` auslösen. Sie können `document.addEventListener()` verwenden, das als erstes Argument das Ereignis (in Anführungszeichen) und als zweites Argument den Callback entgegennimmt.

Entfernen Sie dann in `componentWillUnmount()` denselben Ereignis-Listener. Sie können dieselben Argumente an `document.removeEventListener()` übergeben. Es ist gute Praxis, diese Lifecycle-Methode zu verwenden, um vor dem Unmounten und Zerstören von React-Komponenten Aufräumarbeiten durchzuführen. Das Entfernen von Ereignis-Listenern ist ein Beispiel für eine solche Aufräumaktion.

# --hints--

`MyComponent` sollte ein `div`-Element rendern, das ein `h1`-Tag umschließt.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

Ein `keydown`-Listener sollte in `componentDidMount` an das document angehängt werden.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

Der `keydown`-Listener sollte in `componentWillUnmount` vom document entfernt werden.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

Sobald die Komponente gemountet ist, sollte das Drücken von `enter` ihren Zustand und das gerenderte `h1`-Tag aktualisieren.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
