---
id: 5a24c314108439a4d403618b
title: Geschwisterelementen ein eindeutiges Key-Attribut geben
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

Die letzte Aufgabe zeigte, wie die `map`-Methode verwendet wird, um dynamisch eine Anzahl von Elementen basierend auf Benutzereingaben zu rendern. Allerdings fehlte in diesem Beispiel ein wichtiger Teil. Wenn Sie ein Array von Elementen erstellen, muss jedes ein `key`-Attribut mit einem eindeutigen Wert erhalten. React verwendet diese Keys, um nachzuverfolgen, welche Elemente hinzugefügt, geändert oder entfernt werden. Das hilft, den Prozess des erneuten Renderns effizienter zu gestalten, wenn die Liste auf irgendeine Weise verändert wird.

**Hinweis:** Keys müssen nur zwischen Geschwisterelementen eindeutig sein, sie müssen nicht global in Ihrer Anwendung eindeutig sein.

# --instructions--

Der Code-Editor enthält ein Array mit einigen Frontend-Frameworks und eine zustandslose funktionale Komponente namens `Frameworks()`. `Frameworks()` soll das Array in eine ungeordnete Liste umwandeln, ähnlich wie in der letzten Aufgabe. Beenden Sie die `map`-Callback-Funktion, um für jedes Framework im `li`-Array ein `frontEndFrameworks`-Element zurückzugeben. Diesmal stellen Sie sicher, dass jedes `li` ein `key`-Attribut erhält, das auf einen eindeutigen Wert gesetzt ist. Die `li`-Elemente sollten außerdem Text aus `frontEndFrameworks` enthalten.

Normalerweise sollten Sie den Key so wählen, dass er das gerenderte Element eindeutig identifiziert. Im äußersten Notfall kann der Array-Index verwendet werden, aber in der Regel sollten Sie versuchen, eine eindeutige Identifikation zu verwenden.

# --hints--

Die `Frameworks`-Komponente sollte existieren und auf der Seite gerendert werden.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` sollte ein `h1`-Element rendern.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` sollte ein `ul`-Element rendern.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

Das `ul`-Tag sollte 6 untergeordnete `li`-Elemente rendern.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

Jedes Listenelement sollte ein eindeutiges `key`-Attribut haben.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

Jedes Listenelement sollte Text aus `frontEndFrameworks` enthalten.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
