---
id: 5a24c314108439a4d403616c
title: Standard-Props überschreiben
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

Die Möglichkeit, Standard-Props zu setzen, ist eine nützliche Funktion in React. Um die Standard-Props zu überschreiben, setzen Sie die Prop-Werte für eine Komponente explizit.

# --instructions--

Die `ShoppingCart`-Komponente rendert jetzt eine Kindkomponente `Items`. Diese `Items`-Komponente hat eine Standard-Prop `quantity`, die auf die Ganzzahl `0` gesetzt ist. Überschreiben Sie die Standard-Prop, indem Sie einen Wert von `10` für `quantity` übergeben.

**Hinweis:** Denken Sie daran, dass die Syntax zum Hinzufügen einer Prop zu einer Komponente ähnlich aussieht wie das Hinzufügen von HTML-Attributen. Da der Wert für `quantity` jedoch eine Ganzzahl ist, wird er nicht in Anführungszeichen gesetzt, sondern sollte in geschweifte Klammern eingeschlossen werden. Zum Beispiel `{100}`. Diese Syntax weist JSX an, den Wert innerhalb der Klammern direkt als JavaScript zu interpretieren.

# --hints--

Die Komponente `ShoppingCart` sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Die Komponente `Items` sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Die `Items`-Komponente sollte eine Prop `{ quantity: 10 }` erhalten, die von der `ShoppingCart`-Komponente übergeben wird.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
