---
id: 5a24c314108439a4d4036167
title: Eine Klassenkomponente im DOM rendern
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

Sie erinnern sich vielleicht daran, dass Sie in einer früheren Aufgabe die ReactDOM-API verwendet haben, um JSX-Elemente im DOM zu rendern. Der Prozess zum Rendern von React-Komponenten sieht sehr ähnlich aus. Die letzten Aufgaben konzentrierten sich auf Komponenten und Komposition, daher wurde das Rendern für Sie im Hintergrund erledigt. Allerdings wird kein React-Code, den Sie schreiben, im DOM gerendert, ohne dass Sie die ReactDOM-API aufrufen.

Hier eine Auffrischung der Syntax: `ReactDOM.render(componentToRender, targetNode)`. Das erste Argument ist die React-Komponente, die Sie rendern möchten. Das zweite Argument ist der DOM-Knoten, in dem Sie diese Komponente rendern möchten.

React-Komponenten werden in `ReactDOM.render()` etwas anders übergeben als JSX-Elemente. Bei JSX-Elementen übergeben Sie den Namen des Elements, das Sie rendern möchten. Bei React-Komponenten müssen Sie jedoch dieselbe Syntax verwenden, als ob Sie eine verschachtelte Komponente rendern würden, zum Beispiel `ReactDOM.render(<ComponentToRender />, targetNode)`. Diese Syntax verwenden Sie sowohl für ES6-Klassenkomponenten als auch für funktionale Komponenten.

# --instructions--

Sowohl die `Fruits` als auch die `Vegetables` Komponenten sind für Sie im Hintergrund definiert. Rendern Sie beide Komponenten als Kinder der `TypesOfFood` Komponente, und rendern Sie dann `TypesOfFood` ins DOM. Es steht Ihnen ein `div` mit `id='challenge-node'` zur Verfügung.

# --hints--

Die `TypesOfFood` Komponente sollte ein einzelnes `div` Element zurückgeben.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Die `TypesOfFood` Komponente sollte die `Fruits` Komponente nach dem `h1` Element rendern.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

Die `TypesOfFood` Komponente sollte die `Vegetables` Komponente nach `Fruits` rendern.

```js
assert(
  (function () {
    globalThis.Fruits = function Fruits() {
      return React.createElement('div', null);
    };
    globalThis.Vegetables = function Vegetables() {
      return React.createElement('div', null);
    };
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

Die `TypesOfFood` Komponente sollte im DOM innerhalb des `div` mit der ID `challenge-node` gerendert werden.

```js
assert(
  (function () {
    const hasRenderCall = /ReactDOM\.render\(<TypesOfFood\/>,document\.getElementById\((['"`])challenge-node\1\)\);?/.test(
      __helpers.removeWhiteSpace(code)
    );
    if (!hasRenderCall) {
      return false;
    }

    const e = React.createElement;
    globalThis.Fruits = function Fruits() {
      return e(
        'div',
        null,
        e('h2', null, 'Fruits:'),
        e('h4', null, 'Non-Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Apples'),
          e('li', null, 'Blueberries'),
          e('li', null, 'Strawberries'),
          e('li', null, 'Bananas')
        ),
        e('h4', null, 'Citrus:'),
        e(
          'ul',
          null,
          e('li', null, 'Lemon'),
          e('li', null, 'Lime'),
          e('li', null, 'Orange'),
          e('li', null, 'Grapefruit')
        )
      );
    };
    globalThis.Vegetables = function Vegetables() {
      return e(
        'div',
        null,
        e('h2', null, 'Vegetables:'),
        e(
          'ul',
          null,
          e('li', null, 'Brussel Sprouts'),
          e('li', null, 'Broccoli'),
          e('li', null, 'Squash')
        )
      );
    };

    ReactDOM.render(
      React.createElement(TypesOfFood),
      document.getElementById('challenge-node')
    );

    const root = document.getElementById('challenge-node').childNodes[0];
    if (!root) {
      return false;
    }
    const html = root.innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
