---
id: 5a24c314108439a4d4036185
title: Verwenden Sie && für eine prägnantere Bedingung
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

Die `if/else`-Anweisungen funktionierten in der letzten Aufgabe, aber es gibt eine prägnantere Möglichkeit, dasselbe Ergebnis zu erzielen. Stellen Sie sich vor, Sie verfolgen mehrere Bedingungen in einer Komponente und möchten, dass je nach jeder dieser Bedingungen unterschiedliche Elemente gerendert werden. Wenn Sie viele `else if`-Anweisungen schreiben, um leicht unterschiedliche Benutzeroberflächen zurückzugeben, wiederholen Sie möglicherweise Code, was Raum für Fehler lässt. Stattdessen können Sie den `&&`-logischen Operator verwenden, um bedingte Logik auf eine prägnantere Weise auszuführen. Das ist möglich, weil Sie prüfen möchten, ob eine Bedingung `true` ist, und wenn ja, etwas Auszeichnung zurückgeben. Hier ist ein Beispiel:

```jsx
{condition && <p>markup</p>}
```

Wenn die `condition` `true` ist, wird die Auszeichnung zurückgegeben. Wenn die Bedingung `false` ist, gibt die Operation sofort `false` zurück, nachdem sie die `condition` ausgewertet hat, und gibt nichts zurück. Sie können diese Anweisungen direkt in Ihrem JSX einfügen und mehrere Bedingungen verketten, indem Sie nach jeder `&&` schreiben. So können Sie komplexere bedingte Logik in Ihrer `render()`-Methode handhaben, ohne viel Code zu wiederholen.

# --instructions--

Lösen Sie das vorherige Beispiel erneut, sodass die `h1` nur gerendert wird, wenn `display` `true` ist, verwenden Sie jedoch den `&&`-logischen Operator anstelle einer `if/else`-Anweisung.

# --hints--

`MyComponent` sollte existieren und gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

Wenn `display` auf `true` gesetzt ist, sollten ein `div`, `button` und `h1` gerendert werden.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

Wenn `display` auf `false` gesetzt ist, sollten nur ein `div` und `button` gerendert werden.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

Die render-Methode sollte den `&&`-logischen Operator verwenden, um die Bedingung von `this.state.display` zu prüfen.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
