---
id: 5a24c314108439a4d403616d
title: "Verwenden Sie PropTypes, um die erwarteten Props zu definieren"
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React bietet nützliche Typprüfungsfunktionen, um zu überprüfen, ob Komponenten Props des richtigen Typs erhalten. Zum Beispiel ruft Ihre Anwendung eine API auf, um Daten abzurufen, von denen Sie erwarten, dass sie in einem Array vorliegen, das dann als Prop an eine Komponente übergeben wird. Sie können `propTypes` in Ihrer Komponente festlegen, um zu verlangen, dass die Daten vom Typ `array` sind. Dies wirft eine hilfreiche Warnung, wenn die Daten einen anderen Typ haben.

Es gilt als bewährte Vorgehensweise, `propTypes` festzulegen, wenn Sie den Typ einer Prop im Voraus kennen. Sie können eine `propTypes`-Eigenschaft für eine Komponente auf die gleiche Weise definieren, wie Sie `defaultProps` definiert haben. Dadurch wird überprüft, dass Props mit einem bestimmten Schlüssel vorhanden sind und einen bestimmten Typ haben. Hier ist ein Beispiel, um den Typ `function` für eine Prop namens `handleClick` zu verlangen:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

Im obigen Beispiel überprüft der Teil `PropTypes.func`, dass `handleClick` eine Funktion ist. Das Hinzufügen von `isRequired` teilt React mit, dass `handleClick` eine erforderliche Eigenschaft für diese Komponente ist. Sie sehen eine Warnung, wenn diese Prop nicht bereitgestellt wird. Beachten Sie auch, dass `func` `function` darstellt. Unter den sieben JavaScript-Primitivtypen sind `function` und `boolean` (geschrieben als `bool`) die einzigen beiden mit ungewöhnlicher Schreibweise. Zusätzlich zu den Primitivtypen gibt es weitere verfügbare Typen. Zum Beispiel können Sie überprüfen, ob eine Prop ein React-Element ist. Bitte beachten Sie die <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">Dokumentation</a> für alle Optionen.

**Hinweis:** Ab React v15.5.0 wird `PropTypes` unabhängig von React importiert, so: `import PropTypes from 'prop-types';`

# --instructions--

Definieren Sie `propTypes` für die `Items`-Komponente, um `quantity` als Prop zu verlangen und zu überprüfen, dass es vom Typ `number` ist.

# --hints--

Die `ShoppingCart`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Die `Items`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Die `Items`-Komponente sollte eine `propTypes`-Prüfung enthalten, um einen Wert für `quantity` zu verlangen und sicherzustellen, dass dessen Wert eine Zahl ist.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --before-all--

```js
var PropTypes = {
  number: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
