---
id: 5a24c314108439a4d403616b
title: Standard-Props verwenden
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

React bietet auch die Möglichkeit, Standard-Props festzulegen. Sie können einer Komponente Standard-Props als Eigenschaft direkt an der Komponente zuweisen, und React weist diese Standard-Prop zu, falls erforderlich. So können Sie angeben, welchen Wert eine Prop haben soll, wenn kein Wert explizit übergeben wird. Wenn Sie zum Beispiel `MyComponent.defaultProps = { location: 'San Francisco' }` deklarieren, haben Sie eine location-Prop definiert, die auf den String `San Francisco` gesetzt ist, sofern Sie nichts anderes angeben. React weist Standard-Props zu, wenn Props undefiniert sind, aber wenn Sie `null` als Wert für eine Prop übergeben, bleibt dieser Wert `null`.

# --instructions--

Der Code-Editor zeigt eine `ShoppingCart`-Komponente. Definieren Sie Standard-Props für diese Komponente, die eine Prop `items` mit dem Wert `0` festlegen.

# --hints--

Die `ShoppingCart`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Die `ShoppingCart`-Komponente sollte eine Standard-Prop von `{ items: 0 }` haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
