---
id: 5a24c314108439a4d4036174
title: "Binden von `this` an eine Klassenmethode"
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

Zusätzlich zum Setzen und Aktualisieren von `state` können Sie auch Methoden für Ihre Komponentenklasse definieren. Eine Klassenmethode muss typischerweise das Schlüsselwort `this` verwenden, damit sie innerhalb des Gültigkeitsbereichs der Methode auf Eigenschaften der Klasse (wie `state` und `props`) zugreifen kann. Es gibt einige Möglichkeiten, damit Ihre Klassenmethoden auf `this` zugreifen können.

Eine gängige Methode ist, `this` explizit im Konstruktor zu binden, sodass `this` an die Klassenmethoden gebunden wird, wenn die Komponente initialisiert wird. Ihnen ist vielleicht aufgefallen, dass die letzte Aufgabe `this.handleClick = this.handleClick.bind(this)` für ihre `handleClick`-Methode im Konstruktor verwendet hat. Wenn Sie dann eine Funktion wie `this.setState()` innerhalb Ihrer Klassenmethode aufrufen, bezieht sich `this` auf die Klasse und ist nicht `undefined`.

**Hinweis:** Das Schlüsselwort `this` ist einer der verwirrendsten Aspekte von JavaScript, spielt aber eine wichtige Rolle in React. Obwohl sein Verhalten hier völlig normal ist, sind diese Lektionen nicht der richtige Ort für eine ausführliche Überprüfung von `this`. Bitte sehen Sie sich andere Lektionen an, falls das oben Verstandene verwirrend ist!

# --instructions--

Der Code-Editor enthält eine Komponente mit einem `state`, das den Text verfolgt. Sie hat auch eine Methode, mit der Sie den Text auf `You clicked!` setzen können. Die Methode funktioniert jedoch nicht, weil sie das undefinierte Schlüsselwort `this` verwendet. Beheben Sie das, indem Sie `this` explizit an die `handleClick()`-Methode im Konstruktor der Komponente binden.

Fügen Sie anschließend einen Klick-Handler zum `button`-Element in der Render-Methode hinzu. Er soll die `handleClick()`-Methode auslösen, wenn der Button ein Klick-Ereignis erhält. Denken Sie daran, dass die Methode, die Sie an den `onClick`-Handler übergeben, geschweifte Klammern benötigt, da sie direkt als JavaScript interpretiert werden soll.

Sobald Sie die oben genannten Schritte abgeschlossen haben, sollten Sie den Button klicken können und `You clicked!` sehen.

# --hints--

`MyComponent` sollte ein `div`-Element zurückgeben, das zwei Elemente umschließt: einen Button und ein `h1`-Element, in dieser Reihenfolge.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

Der Zustand von `MyComponent` sollte mit dem Schlüssel-Wert-Paar `{ text: "Hello" }` initialisiert werden.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

Ein Klick auf das `button`-Element sollte die `handleClick`-Methode ausführen und den Zustand `text` auf `You clicked!` setzen.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
