---
id: 5a24c314108439a4d4036165
title: "React verwenden, um verschachtelte Komponenten zu rendern"
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

Die letzte Aufgabe zeigte eine einfache Möglichkeit, zwei Komponenten zu kombinieren, aber es gibt viele verschiedene Wege, wie Sie Komponenten mit React zusammensetzen können.

Die Komponentenzusammensetzung ist eine der leistungsstarken Funktionen von React. Wenn Sie mit React arbeiten, ist es wichtig, Ihre Benutzeroberfläche in Form von Komponenten zu betrachten, wie im App-Beispiel der letzten Aufgabe. Sie zerlegen Ihre UI in ihre grundlegenden Bausteine, und diese Teile werden zu Komponenten. Das hilft dabei, den Code, der für die Benutzeroberfläche verantwortlich ist, vom Code zu trennen, der die Logik Ihrer Anwendung steuert. Das kann die Entwicklung und Wartung komplexer Projekte erheblich vereinfachen.

# --instructions--

Im Code-Editor sind zwei funktionale Komponenten definiert, die `TypesOfFruit` und `Fruits` heißen. Nehmen Sie die `TypesOfFruit`-Komponente und setzen Sie sie zusammen, oder *verschachteln* Sie sie, innerhalb der `Fruits`-Komponente. Dann nehmen Sie die `Fruits`-Komponente und verschachteln Sie sie innerhalb der `TypesOfFood`-Komponente. Das Ergebnis sollte eine Kindkomponente sein, die innerhalb einer Elternkomponente verschachtelt ist, welche wiederum in einer eigenen Elternkomponente verschachtelt ist!

# --hints--

Die `TypesOfFood`-Komponente sollte ein einzelnes `div`-Element zurückgeben.

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

Die `TypesOfFood`-Komponente sollte die `Fruits`-Komponente zurückgeben.

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

Die `Fruits`-Komponente sollte die `TypesOfFruit`-Komponente zurückgeben.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

Die `TypesOfFruit`-Komponente sollte die `h2`- und `ul`-Elemente zurückgeben.

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
