---
id: 5a24c314108439a4d4036178
title: Erstellen Sie eine kontrollierte Eingabe
challengeType: 6
forumTopicId: 301385
dashedName: create-a-controlled-input
---

# --description--

Ihre Anwendung kann komplexere Interaktionen zwischen `state` und der gerenderten Benutzeroberfläche haben. Zum Beispiel behalten Formularsteuerelemente für Texteingaben, wie `input` und `textarea`, ihren eigenen Zustand im DOM, während der Benutzer tippt. Mit React können Sie diesen veränderlichen Zustand in den `state` einer React-Komponente verschieben. Die Eingabe des Benutzers wird Teil des Anwendungs-`state`, sodass React den Wert dieses Eingabefelds steuert. Typischerweise handelt es sich bei React-Komponenten mit Eingabefeldern, in die der Benutzer tippen kann, um ein kontrolliertes Eingabeformular.

# --instructions--

Der Code-Editor enthält das Grundgerüst einer Komponente namens `ControlledInput`, um ein kontrolliertes `input`-Element zu erstellen. Der `state` der Komponente ist bereits mit einer `input`-Eigenschaft initialisiert, die einen leeren String enthält. Dieser Wert stellt den Text dar, den ein Benutzer in das `input`-Feld eingibt.

Erstellen Sie zuerst eine Methode namens `handleChange()`, die einen Parameter namens `event` hat. Wenn die Methode aufgerufen wird, erhält sie ein `event`-Objekt, das einen String-Text vom `input`-Element enthält. Sie können diesen String innerhalb der Methode mit `event.target.value` zugreifen. Aktualisieren Sie die `input`-Eigenschaft des `state` der Komponente mit diesem neuen String.

Erstellen Sie in der `render`-Methode das `input`-Element oberhalb des `h4`-Tags. Fügen Sie ein `value`-Attribut hinzu, das gleich der `input`-Eigenschaft des `state` der Komponente ist. Fügen Sie dann eine `onChange`-Eigenschaft hinzu, die auf die `handleChange()`-Ereignis-Handler-Methode gesetzt ist.

Wenn Sie in das Eingabefeld tippen, wird dieser Text von der `handleChange()`-Methode verarbeitet, als `input`-Eigenschaft im lokalen `state` gesetzt und als Wert im `input`-Feld auf der Seite gerendert. Die Komponente `state` ist die einzige Quelle der Wahrheit bezüglich der Eingabedaten.

Zuletzt vergessen Sie nicht, die notwendigen Bindungen im Konstruktor hinzuzufügen.

# --hints--

`ControlledInput` sollte ein `div`-Element zurückgeben, das ein `input` und ein `p`-Tag enthält.

```js
assert(
  Enzyme.mount(React.createElement(ControlledInput))
    .find('div')
    .children()
    .find('input').length === 1 &&
    Enzyme.mount(React.createElement(ControlledInput))
      .find('div')
      .children()
      .find('p').length === 1
);
```

Der Zustand von `ControlledInput` sollte mit einer `input`-Eigenschaft initialisiert werden, die auf einen leeren String gesetzt ist.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(ControlledInput)).state('input'),
  ''
);
```

Das Tippen im Eingabeelement sollte den Zustand und den Wert der Eingabe aktualisieren, und das `p`-Element sollte diesen Zustand beim Tippen rendern.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(ControlledInput));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return waitForIt(() => mockedComponent.state('input'));
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => ({
      state: mockedComponent.state('input'),
      text: mockedComponent.find('p').text(),
      inputVal: mockedComponent.find('input').props().value
    }));
  };
  const before = await _1();
  const after = await _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.text === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
```

# --seed--

## --seed-contents--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        { /* Change code below this line */}

        { /* Change code above this line */}
        <h4>Controlled Input:</h4>
        <p>{this.state.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class ControlledInput extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  render() {
    return (
      <div>
        <input
          value={this.state.input}
          onChange={this.handleChange} />
        <h4>Controlled Input:</h4>

        <p>{this.state.input}</p>
      </div>
    );
  }
};
```
