---
id: 5a24c314108439a4d403617c
title: Die Lifecycle-Methode componentWillMount verwenden
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React-Komponenten haben mehrere spezielle Methoden, die Möglichkeiten bieten, Aktionen zu bestimmten Zeitpunkten im Lebenszyklus einer Komponente auszuführen. Diese werden Lifecycle-Methoden oder Lifecycle-Hooks genannt und erlauben es Ihnen, Komponenten zu bestimmten Zeitpunkten abzufangen. Das kann vor dem Rendern, vor dem Aktualisieren, bevor sie Props erhalten, bevor sie unmounten und so weiter sein. Hier ist eine Liste einiger der wichtigsten Lifecycle-Methoden: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` Die nächsten Lektionen behandeln einige der grundlegenden Anwendungsfälle für diese Lifecycle-Methoden.

**Hinweis:** Die `componentWillMount` Lifecycle-Methode wird in einer zukünftigen Version von 16.X veraltet sein und in Version 17 entfernt werden. Erfahren Sie mehr in diesem <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">Artikel</a>

# --instructions--

Die `componentWillMount()`-Methode wird vor der `render()`-Methode aufgerufen, wenn eine Komponente in das DOM eingebunden wird. Protokollieren Sie etwas in der Konsole innerhalb von `componentWillMount()` – Sie sollten die Browser-Konsole geöffnet haben, um die Ausgabe zu sehen.

# --hints--

`MyComponent` sollte ein `div`-Element rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log` sollte in `componentWillMount` aufgerufen werden.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
