---
id: 5a24c314108439a4d403617d
title: Die Lifecycle-Methode componentDidMount verwenden
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

Die meisten Webentwickler müssen irgendwann eine API-Endpunkt aufrufen, um Daten abzurufen. Wenn Sie mit React arbeiten, ist es wichtig zu wissen, wo Sie diese Aktion ausführen.

Die bewährte Vorgehensweise bei React ist, API-Aufrufe oder andere Aufrufe an Ihren Server in der Lifecycle-Methode `componentDidMount()` zu platzieren. Diese Methode wird aufgerufen, nachdem eine Komponente im DOM eingebunden wurde. Alle Aufrufe von `setState()` hier lösen ein erneutes Rendern Ihrer Komponente aus. Wenn Sie in dieser Methode eine API aufrufen und den State mit den von der API zurückgegebenen Daten setzen, wird automatisch ein Update ausgelöst, sobald Sie die Daten erhalten.

# --instructions--

Es gibt einen simulierten API-Aufruf in `componentDidMount()`. Er setzt den State nach 2,5 Sekunden, um das Abrufen von Daten vom Server zu simulieren. Dieses Beispiel fordert die aktuelle Gesamtzahl der aktiven Nutzer für eine Website an. Rendern Sie im render-Methodenblock den Wert von `activeUsers` im `h1` nach dem Text `Active Users:`. Beobachten Sie, was in der Vorschau passiert, und ändern Sie gerne das Timeout, um die verschiedenen Effekte zu sehen.

# --hints--

`MyComponent` sollte ein `div`-Element rendern, das einen `h1`-Tag umschließt.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Der Komponentenzustand sollte mit einer Timeout-Funktion in `componentDidMount` aktualisiert werden.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

Der `h1`-Tag sollte den `activeUsers`-Wert aus dem State von `MyComponent` rendern.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
