---
id: 5a24c314108439a4d4036164
title: Eine Komponente mit Komposition erstellen
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

Nun sehen wir uns an, wie wir mehrere React-Komponenten miteinander kombinieren können. Stellen Sie sich vor, Sie bauen eine App und haben drei Komponenten erstellt: eine `Navbar`, `Dashboard` und `Footer`.

Um diese Komponenten zusammenzusetzen, könnten Sie eine `App` *übergeordnete* Komponente erstellen, die jede dieser drei Komponenten als *Kinder* rendert. Um eine Komponente als Kind in einer React-Komponente zu rendern, fügen Sie den Komponenten-Namen als benutzerdefinierten HTML-Tag in das JSX ein. Zum Beispiel könnten Sie in der `render`-Methode Folgendes schreiben:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Wenn React auf einen benutzerdefinierten HTML-Tag stößt, der auf eine andere Komponente verweist (einen Komponenten-Namen, der in `< />` eingeschlossen ist, wie in diesem Beispiel), rendert es die Auszeichnung für diese Komponente an der Stelle des Tags. Dies veranschaulicht die Eltern-Kind-Beziehung zwischen der `App`-Komponente und den `Navbar`, `Dashboard` und `Footer`.

# --instructions--

Im Code-Editor gibt es eine einfache funktionale Komponente namens `ChildComponent` und eine Klassenkomponente namens `ParentComponent`. Setzen Sie die beiden zusammen, indem Sie die `ChildComponent` innerhalb der `ParentComponent` rendern. Achten Sie darauf, das `ChildComponent`-Tag mit einem Schrägstrich zu schließen.

**Hinweis:** `ChildComponent` ist mit einer ES6-Pfeilfunktion definiert, da dies eine sehr gängige Praxis bei der Verwendung von React ist.

# --hints--

Die React-Komponente sollte ein einzelnes `div`-Element zurückgeben.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

Die Komponente sollte zwei verschachtelte Elemente zurückgeben.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

Die Komponente sollte die `ChildComponent` als ihr zweites Kind zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
