---
id: 5a24c314108439a4d4036177
title: Einen einfachen Zähler schreiben
challengeType: 6
forumTopicId: 301425
dashedName: write-a-simple-counter
---

# --description--

Sie können eine komplexere zustandsbehaftete Komponente entwerfen, indem Sie die bisher behandelten Konzepte kombinieren. Dazu gehören das Initialisieren von `state`, das Schreiben von Methoden, die `state` setzen, und das Zuweisen von Klick-Handlern, um diese Methoden auszulösen.

# --instructions--

Die `Counter`-Komponente verfolgt einen `count`-Wert im `state`. Es gibt zwei Buttons, die die Methoden `increment()` und `decrement()` aufrufen. Schreiben Sie diese Methoden so, dass der Zählerwert um 1 erhöht oder verringert wird, wenn der entsprechende Button geklickt wird. Erstellen Sie außerdem eine `reset()`-Methode, sodass beim Klicken des Reset-Buttons der Zähler auf 0 gesetzt wird.

**Hinweis:** Achten Sie darauf, die `className`s der Buttons nicht zu verändern. Denken Sie auch daran, die notwendigen Bindungen für die neu erstellten Methoden im Konstruktor hinzuzufügen.

# --hints--

`Counter` sollte ein `div`-Element zurückgeben, das drei Buttons mit folgendem Textinhalt in dieser Reihenfolge enthält: `Increment!`, `Decrement!`, `Reset`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Counter));
    return (
      mockedComponent.find('.inc').text() === 'Increment!' &&
      mockedComponent.find('.dec').text() === 'Decrement!' &&
      mockedComponent.find('.reset').text() === 'Reset'
    );
  })()
);
```

Der Zustand von `Counter` sollte mit einer `count`-Eigenschaft initialisiert werden, die auf `0` gesetzt ist.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
assert(mockedComponent.find('h1').text() === 'Current Count: 0');
```

Ein Klick auf den Erhöhen-Button sollte den Zähler um `1` erhöhen.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.inc').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: 1');
```

Ein Klick auf den Verringern-Button sollte den Zähler um `1` verringern.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.find('.dec').simulate('click');
assert(mockedComponent.find('h1').text() === 'Current Count: -1');
```

Ein Klick auf den Reset-Button sollte den Zähler auf `0` zurücksetzen.

```js
const mockedComponent = Enzyme.mount(React.createElement(Counter));
mockedComponent.setState({ count: 5 });
const currentCountElement = mockedComponent.find('h1');
assert(currentCountElement.text() === 'Current Count: 5');
mockedComponent.find('.reset').simulate('click');
assert(currentCountElement.text() === 'Current Count: 0');
```

# --seed--

## --seed-contents--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class Counter extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      count: 0
    };
  this.increment = this.increment.bind(this);
 this.decrement = this.decrement.bind(this);
 this.reset = this.reset.bind(this);
 }
  reset() {
    this.setState({
      count: 0
    });
  }
  increment() {
    this.setState(state => ({
      count: state.count + 1
    }));
  }
  decrement() {
    this.setState(state => ({
      count: state.count - 1
    }));
  }
  render() {
    return (
      <div>
        <button className='inc' onClick={this.increment}>Increment!</button>
        <button className='dec' onClick={this.decrement}>Decrement!</button>
        <button className='reset' onClick={this.reset}>Reset</button>
        <h1>Current Count: {this.state.count}</h1>
      </div>
    );
  }
};
```
