---
id: 5a24c314108439a4d4036169
title: Props an eine zustandslose funktionale Komponente übergeben
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

Die vorherigen Aufgaben behandelten viel über das Erstellen und Zusammensetzen von JSX-Elementen, funktionalen Komponenten und ES6-Klassenkomponenten in React. Mit dieser Grundlage ist es Zeit, sich eine weitere in React sehr häufig verwendete Funktion anzusehen: **Props**. In React können Sie Props, also Eigenschaften, an Kindkomponenten übergeben. Angenommen, Sie haben eine `App`-Komponente, die eine Kindkomponente namens `Welcome` rendert, welche eine zustandslose funktionale Komponente ist. Sie können `Welcome` eine `user`-Eigenschaft übergeben, indem Sie schreiben:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

Sie verwenden **benutzerdefinierte HTML-Attribute**, die Sie erstellen und die von React unterstützt werden, um sie an die Komponente zu übergeben. In diesem Fall wird die erstellte Eigenschaft `user` an die Komponente `Welcome` übergeben. Da `Welcome` eine zustandslose funktionale Komponente ist, hat sie Zugriff auf diesen Wert wie folgt:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

Es ist üblich, diesen Wert `props` zu nennen, und bei zustandslosen funktionalen Komponenten betrachten Sie ihn im Grunde als Argument einer Funktion, die JSX zurückgibt. Sie können den Wert des Arguments im Funktionskörper abrufen. Bei Klassenkomponenten ist das etwas anders.

# --instructions--

Im Code-Editor befinden sich `Calendar` und `CurrentDate`-Komponenten. Wenn Sie `CurrentDate` aus der `Calendar`-Komponente rendern, übergeben Sie eine Eigenschaft `date`, der das aktuelle Datum aus dem `Date`-Objekt von JavaScript zugewiesen ist. Greifen Sie dann in der `prop`-Komponente auf diese `CurrentDate` zu und zeigen Sie ihren Wert innerhalb der `p`-Tags an. Beachten Sie, dass `prop`-Werte als JavaScript ausgewertet werden müssen, wenn sie in geschweifte Klammern eingeschlossen sind, zum Beispiel `date={Date()}`.

# --hints--

Die `Calendar`-Komponente sollte ein einzelnes `div`-Element zurückgeben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Das zweite Kind der `Calendar`-Komponente sollte die `CurrentDate`-Komponente sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

Die `CurrentDate`-Komponente sollte eine Prop namens `date` haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

Die `date`-Prop der `CurrentDate` sollte einen Textstring enthalten.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

Die `date`-Prop sollte durch Aufrufen von `Date()` erzeugt werden.

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

Die `CurrentDate`-Komponente sollte den Wert der `date`-Prop im `p`-Tag rendern.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
