---
id: 5a24c314108439a4d4036170
title: Einen zustandsbehafteten komponenten erstellen
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

Eines der wichtigsten Themen in React ist `state`. Der Zustand besteht aus allen Daten, die Ihre Anwendung kennen muss und die sich im Laufe der Zeit ändern können. Sie möchten, dass Ihre Apps auf Zustandsänderungen reagieren und bei Bedarf eine aktualisierte Benutzeroberfläche anzeigen. React bietet eine gute Lösung für das Zustandsmanagement moderner Webanwendungen.

Sie erstellen Zustand in einer React-Komponente, indem Sie eine `state`-Eigenschaft in der `constructor` der Komponente deklarieren. Dadurch wird die Komponente mit `state` initialisiert, wenn sie erstellt wird. Die `state`-Eigenschaft muss auf einen JavaScript-`object` gesetzt werden. Die Deklaration sieht so aus:

```jsx
this.state = {

}
```

Sie haben während der gesamten Lebensdauer Ihrer Komponente Zugriff auf das `state`-Objekt. Sie können es aktualisieren, in Ihrer Benutzeroberfläche rendern und als Props an untergeordnete Komponenten weitergeben. Das `state`-Objekt kann so komplex oder so einfach sein, wie Sie es benötigen. Beachten Sie, dass Sie eine Klassenkomponente erstellen müssen, indem Sie `React.Component` erweitern, um `state` auf diese Weise zu erstellen.

# --instructions--

Im Code-Editor befindet sich eine Komponente, die versucht, eine `firstName`-Eigenschaft aus ihrem `state` zu rendern. Es ist jedoch kein `state` definiert. Initialisieren Sie die Komponente mit `state` in der `constructor` und weisen Sie eine Eigenschaft von `firstName` Ihren Namen zu.

# --hints--

`StatefulComponent` sollte existieren und rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent` sollte ein `div`- und ein `h1`-Element rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Der Zustand von `StatefulComponent` sollte mit einer Eigenschaft `firstName` initialisiert werden, die auf einen String gesetzt ist.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

Die Eigenschaft `firstName` im Zustand von `StatefulComponent` sollte im `h1`-Element gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
