---
id: 5a24c314108439a4d4036189
title: Inline-CSS bedingt basierend auf dem Zustand der Komponente ändern
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

Bis jetzt haben Sie mehrere Anwendungen von bedingtem Rendern und die Verwendung von inline Styles gesehen. Hier ist ein weiteres Beispiel, das beide Themen kombiniert. Sie können CSS auch bedingt rendern, basierend auf dem Zustand einer React-Komponente. Dazu prüfen Sie eine Bedingung, und wenn diese erfüllt ist, ändern Sie das Styles-Objekt, das den JSX-Elementen in der render-Methode zugewiesen ist.

Dieses Paradigma ist wichtig zu verstehen, weil es eine dramatische Abkehr von der traditionelleren Methode darstellt, Styles durch direkte Änderung von DOM-Elementen anzuwenden (was zum Beispiel bei jQuery sehr verbreitet ist). Bei diesem Ansatz müssen Sie verfolgen, wann sich Elemente ändern, und auch die eigentliche Manipulation direkt vornehmen. Es kann schwierig werden, Änderungen nachzuvollziehen, was Ihre Benutzeroberfläche unvorhersehbar machen kann. Wenn Sie ein Style-Objekt basierend auf einer Bedingung setzen, beschreiben Sie, wie die Benutzeroberfläche als Funktion des Zustands der Anwendung aussehen soll. Es gibt einen klaren Informationsfluss, der nur in eine Richtung verläuft. Dies ist die bevorzugte Methode beim Schreiben von Anwendungen mit React.

# --instructions--

Der Code-Editor enthält eine einfache kontrollierte Eingabekomponente mit einem gestylten Rahmen. Sie möchten diesen Rahmen rot stylen, wenn der Benutzer mehr als 15 Zeichen Text in das Eingabefeld eingibt. Fügen Sie eine Bedingung hinzu, um dies zu prüfen, und wenn die Bedingung erfüllt ist, setzen Sie den Rahmenstil der Eingabe auf `3px solid red`. Sie können es ausprobieren, indem Sie Text in die Eingabe eingeben.

# --hints--

Die `GateKeeper`-Komponente sollte ein `div`-Element rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Die `GateKeeper`-Komponente sollte mit einem State-Schlüssel `input` initialisiert werden, der auf einen leeren String gesetzt ist.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

Die `GateKeeper`-Komponente sollte ein `h3`-Tag und ein `input`-Tag rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

Das `input`-Tag sollte anfangs einen Style von `1px solid black` für die `border`-Eigenschaft haben.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

Das `input`-Tag sollte mit einem Rahmen von `3px solid red` gestylt werden, wenn der Eingabewert im State länger als 15 Zeichen ist.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
