---
id: 5a24c314108439a4d4036173
title: "Setzen von `state` mit this.setState"
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

Die vorherigen Aufgaben behandelten die Komponente `state` und wie man den `constructor` initialisiert. Es gibt auch eine Möglichkeit, den `state` der Komponente zu ändern. React stellt eine Methode zum Aktualisieren des Komponenten-`state` bereit, die `setState` heißt. Sie rufen die `setState`-Methode innerhalb Ihrer Komponentenklasse folgendermaßen auf: `this.setState()`, wobei Sie ein Objekt mit Schlüssel-Wert-Paaren übergeben. Die Schlüssel sind Ihre `state`-Eigenschaften und die Werte sind die aktualisierten `state`-Daten. Wenn wir zum Beispiel einen `username` im `state` speichern und diesen aktualisieren wollten, sähe das so aus:

```jsx
this.setState({
  username: 'Lewis'
});
```

React erwartet, dass Sie den `state` niemals direkt verändern, sondern immer `this.setState()` verwenden, wenn sich der `state` ändert. Außerdem sollten Sie beachten, dass React mehrere `state`-Updates zusammenfassen kann, um die Leistung zu verbessern. Das bedeutet, dass `state`-Updates über die `setState`-Methode asynchron sein können. Es gibt eine alternative Syntax für die `setState`-Methode, die eine Möglichkeit bietet, dieses Problem zu umgehen. Das wird selten benötigt, aber es ist gut, es im Hinterkopf zu behalten! Weitere Details finden Sie in unserem <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">React-Artikel</a>.

# --instructions--

Im Code-Editor gibt es ein `button`-Element mit einem `onClick()`-Handler. Dieser Handler wird ausgelöst, wenn das `button` im Browser ein Klick-Ereignis erhält, und führt die `handleClick`-Methode aus, die auf `MyComponent` definiert ist. Aktualisieren Sie innerhalb der `handleClick`-Methode den Komponenten-`state` mithilfe von `this.setState()`. Setzen Sie die `name`-Eigenschaft in `state` auf den String `React Rocks!`.

Klicken Sie auf die Schaltfläche und beobachten Sie, wie sich der gerenderte `state` aktualisiert. Machen Sie sich keine Sorgen, wenn Sie den Code des Klick-Handlers noch nicht vollständig verstehen. Das wird in den kommenden Aufgaben behandelt.

# --hints--

Der `state` von `MyComponent` sollte mit dem Schlüssel-Wert-Paar `{ name: Initial State }` initialisiert werden.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` sollte ein `h1`-Überschriftenelement rendern.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

Das gerenderte `h1`-Überschriftenelement sollte Text enthalten, der aus dem `state` der Komponente gerendert wird.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

Das Aufrufen der `handleClick`-Methode auf `MyComponent` sollte die Eigenschaft name im `state` auf `React Rocks!` setzen.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
