---
id: 5a24c314108439a4d403618c
title: "Array.filter() verwenden, um ein Array dynamisch zu filtern"
challengeType: 6
forumTopicId: 301416
dashedName: use-array-filter-to-dynamically-filter-an-array
---

# --description--

Die `map` Array-Methode ist ein mächtiges Werkzeug, das Sie beim Arbeiten mit React häufig verwenden werden. Eine weitere Methode, die mit `map` zusammenhängt, ist `filter`, die den Inhalt eines Arrays basierend auf einer Bedingung filtert und dann ein neues Array zurückgibt. Zum Beispiel, wenn Sie ein Array von Benutzern haben, die alle eine Eigenschaft `online` besitzen, die auf `true` oder `false` gesetzt sein kann, können Sie nur die Benutzer filtern, die online sind, indem Sie schreiben:

```js
let onlineUsers = users.filter(user => user.online);
```

# --instructions--

Im Code-Editor wird `MyComponent`s `state` mit einem Array von Benutzern initialisiert. Einige Benutzer sind online, andere nicht. Filtern Sie das Array so, dass Sie nur die Benutzer sehen, die online sind. Verwenden Sie dazu zuerst `filter`, um ein neues Array zurückzugeben, das nur die Benutzer enthält, deren `online`-Eigenschaft auf `true` gesetzt ist. Dann mappen Sie im `renderOnline`-Variablen über das gefilterte Array und geben für jeden Benutzer ein `li`-Element zurück, das den Text ihres `username` enthält. Achten Sie darauf, auch ein einzigartiges `key` einzuschließen, wie in den vorherigen Aufgaben.

# --hints--

`MyComponent` sollte existieren und auf der Seite gerendert werden.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('MyComponent').length,
  1
);
```

`MyComponent`s state sollte mit einem Array von sechs Benutzern initialisiert sein.

```js
assert(
  Array.isArray(
    Enzyme.mount(React.createElement(MyComponent)).state('users')
  ) === true &&
    Enzyme.mount(React.createElement(MyComponent)).state('users').length === 6
);
```

`MyComponent` sollte ein `div`, ein `h1` und dann eine ungeordnete Liste zurückgeben, die `li`-Elemente für jeden Benutzer enthält, dessen Online-Status auf `true` gesetzt ist.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const result = () => comp.find('li').length;
  const _1 = result();
  const _2 = () => {
    comp.setState(users(true));
    return result();
  };
  const _3 = () => {
    comp.setState(users(false));
    return result();
  };
  const _4 = () => {
    comp.setState({ users: [] });
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  assert(
    comp.find('div').length === 1 &&
      comp.find('h1').length === 1 &&
      comp.find('ul').length === 1 &&
      _1 === 4 &&
      _2_val === 5 &&
      _3_val === 0 &&
      _4_val === 0
  );
})();
```

`MyComponent` sollte `li`-Elemente rendern, die den `username` jedes Online-Benutzers enthalten.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MyComponent));
  const users = (bool) => ({
    users: [
      { username: 'Jeff', online: bool },
      { username: 'Alan', online: bool },
      { username: 'Mary', online: bool },
      { username: 'Jim', online: bool },
      { username: 'Laura', online: bool }
    ]
  });
  const ul = () => {
    comp.setState(users(true));
    return comp.find('ul').html();
  };
  const html = ul();
  assert(
    html ===
      '<ul><li>Jeff</li><li>Alan</li><li>Mary</li><li>Jim</li><li>Laura</li></ul>'
  );
})();
```

Jedes Listenelement sollte ein einzigartiges `key`-Attribut haben.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(MyComponent)).find('ul');
    console.log(ul.debug());
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key()
    ]);
    return keys.size === 4;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = null; // Change this line
    const renderOnline = null; // Change this line
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      users: [
        {
          username: 'Jeff',
          online: true
        },
        {
          username: 'Alan',
          online: false
        },
        {
          username: 'Mary',
          online: true
        },
        {
          username: 'Jim',
          online: false
        },
        {
          username: 'Sara',
          online: true
        },
        {
          username: 'Laura',
          online: true
        }
      ]
    };
  }
  render() {
    const usersOnline = this.state.users.filter(user => {
      return user.online;
    });
    const renderOnline = usersOnline.map(user => {
      return <li key={user.username}>{user.username}</li>;
    });
    return (
      <div>
        <h1>Current Online Users:</h1>
        <ul>{renderOnline}</ul>
      </div>
    );
  }
}
```
