---
id: 5a24c314108439a4d403618a
title: "Array.map() verwenden, um Elemente dynamisch zu rendern"
challengeType: 6
forumTopicId: 301417
dashedName: use-array-map-to-dynamically-render-elements
---

# --description--

Bedingtes Rendern ist nützlich, aber manchmal müssen Ihre Komponenten eine unbekannte Anzahl von Elementen rendern. Oft weiß ein Programmierer in der reaktiven Programmierung erst zur Laufzeit, wie der Zustand einer Anwendung ist, weil so viel von der Interaktion eines Benutzers mit dem Programm abhängt. Programmierer müssen ihren Code so schreiben, dass er diesen unbekannten Zustand im Voraus korrekt behandelt. Die Verwendung von `Array.map()` in React veranschaulicht dieses Konzept.

Zum Beispiel erstellen Sie eine einfache „To-Do-Liste“-Anwendung. Als Programmierer wissen Sie nicht, wie viele Elemente ein Benutzer auf seiner Liste haben könnte. Sie müssen Ihre Komponente so einrichten, dass sie dynamisch die richtige Anzahl von Listenelementen rendert, lange bevor jemand, der das Programm benutzt, entscheidet, dass heute Waschtag ist.

# --instructions--

Der Code-Editor enthält den größten Teil der `MyToDoList`-Komponente. Ein Teil dieses Codes sollte Ihnen vertraut vorkommen, wenn Sie die Aufgabe zum kontrollierten Formular abgeschlossen haben. Sie werden ein `textarea` und ein `button` bemerken, zusammen mit einigen Methoden, die deren Zustände verfolgen, aber noch wird nichts auf der Seite gerendert.

Erstellen Sie innerhalb von `constructor` ein `this.state`-Objekt und definieren Sie zwei Zustände: `userInput` sollte als leerer String initialisiert werden, und `toDoList` sollte als leeres Array initialisiert werden. Löschen Sie anschließend den `null`-Wert in der `render()`-Methode neben der `items`-Variablen. Ersetzen Sie ihn durch eine Abbildung (map) über das `toDoList`-Array, das im internen Zustand der Komponente gespeichert ist, und rendern Sie dynamisch ein `li` für jedes Element. Versuchen Sie, den String `eat, code, sleep, repeat` in das `textarea` einzugeben, klicken Sie dann auf den Button und sehen Sie, was passiert.

**Hinweis:** Sie wissen vielleicht, dass alle Geschwister-Kindelemente, die durch eine solche Mapping-Operation erzeugt werden, mit einem eindeutigen `key`-Attribut versehen sein müssen. Machen Sie sich keine Sorgen, das ist Thema der nächsten Aufgabe.

# --hints--

Die MyToDoList-Komponente sollte existieren und auf der Seite gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return mockedComponent.find('MyToDoList').length === 1;
  })()
);
```

Das erste Kindelement von `MyToDoList` sollte ein `textarea`-Element sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(0).type() ===
      'textarea'
    );
  })()
);
```

Das zweite Kindelement von `MyToDoList` sollte ein `br`-Element sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(1).type() === 'br'
    );
  })()
);
```

Das dritte Kindelement von `MyToDoList` sollte ein `button`-Element sein.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    return (
      mockedComponent.find('MyToDoList').children().childAt(2).type() ===
      'button'
    );
  })()
);
```

Der Zustand von `MyToDoList` sollte mit `toDoList` als leerem Array initialisiert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      Array.isArray(initialState.toDoList) === true &&
      initialState.toDoList.length === 0
    );
  })()
);
```

Der Zustand von `MyToDoList` sollte mit `userInput` als leerem String initialisiert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
    const initialState = mockedComponent.state();
    return (
      typeof initialState.userInput === 'string' &&
      initialState.userInput.length === 0
    );
  })()
);
```

Wenn der `Create List`-Button geklickt wird, sollte die `MyToDoList`-Komponente dynamisch eine ungeordnete Liste zurückgeben, die für jedes Element einer durch Kommas getrennten Liste, die in das `textarea`-Element eingegeben wurde, ein Listenelement enthält.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyToDoList));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  const state_1 = () => {
    return mockedComponent.find('ul').find('li');
  };
  const setInput = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      'testA, testB, testC'
    );
  };
  const click = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_2 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const setInput_2 = () => {
    return simulateChange(
      mockedComponent.find('textarea'),
      't1, t2, t3, t4, t5, t6'
    );
  };
  const click_1 = () => {
    return mockedComponent.find('button').simulate('click');
  };
  const state_3 = () => {
    const nodes = mockedComponent.find('ul').find('li');
    return { nodes, text: nodes.reduce((t, n) => t + n.text().trim(), '') };
  };
  const awaited_state_1 = state_1();
  const awaited_setInput = setInput();
  const awaited_click = click();
  const awaited_state_2 = state_2();
  const awaited_setInput_2 = setInput_2();
  const awaited_click_1 = click_1();
  const awaited_state_3 = state_3();
  assert(
    awaited_state_1.length === 0 &&
      awaited_state_2.nodes.length === 3 &&
      awaited_state_3.nodes.length === 6 &&
      awaited_state_2.text === 'testAtestBtestC' &&
      awaited_state_3.text === 't1t2t3t4t5t6'
  );
})();
```

# --seed--

## --seed-contents--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    // Change code below this line

    // Change code above this line
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = null; // Change this line
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const textAreaStyles = {
  width: 235,
  margin: 5
};

class MyToDoList extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      toDoList: [],
      userInput: ''
    };
    this.handleSubmit = this.handleSubmit.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  handleSubmit() {
    const itemsArray = this.state.userInput.split(',');
    this.setState({
      toDoList: itemsArray
    });
  }
  handleChange(e) {
    this.setState({
      userInput: e.target.value
    });
  }
  render() {
    const items = this.state.toDoList.map((item, i) => {
      return <li key={i}>{item}</li>;
    });
    return (
      <div>
        <textarea
          onChange={this.handleChange}
          value={this.state.userInput}
          style={textAreaStyles}
          placeholder='Separate Items With Commas'
        />
        <br />
        <button onClick={this.handleSubmit}>Create List</button>
        <h1>My "To Do" List:</h1>
        <ul>{items}</ul>
      </div>
    );
  }
}
```
