---
id: 5a24c314108439a4d4036188
title: Bedingtes Rendern anhand von Props
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

Bis jetzt haben Sie gesehen, wie man `if/else`, `&&` und den ternären Operator (`condition ? expressionIfTrue : expressionIfFalse`) verwendet, um bedingte Entscheidungen darüber zu treffen, was und wann gerendert wird. Es gibt jedoch ein wichtiges Thema, das noch zu besprechen ist und das es Ihnen ermöglicht, alle oder einige dieser Konzepte mit einer weiteren mächtigen React-Funktion zu kombinieren: Props. Die Verwendung von Props zum bedingten Rendern von Code ist bei React-Entwicklern sehr verbreitet – das heißt, sie nutzen den Wert einer bestimmten Prop, um automatisch Entscheidungen darüber zu treffen, was gerendert wird.

In dieser Aufgabe richten Sie eine Kind-Komponente ein, die Rendering-Entscheidungen basierend auf Props trifft. Sie verwenden auch den ternären Operator, aber Sie können sehen, wie mehrere der anderen Konzepte, die in den letzten Aufgaben behandelt wurden, in diesem Zusammenhang ebenso nützlich sein könnten.

# --instructions--

Der Code-Editor enthält zwei Komponenten, die teilweise für Sie definiert sind: eine Elternkomponente namens `GameOfChance` und eine Kind-Komponente namens `Results`. Sie werden verwendet, um ein einfaches Spiel zu erstellen, bei dem der Benutzer eine Taste drückt, um zu sehen, ob er gewinnt oder verliert.

Zuerst benötigen Sie einen einfachen Ausdruck, der bei jedem Ausführen zufällig einen anderen Wert zurückgibt. Sie können `Math.random()` verwenden. Diese Methode gibt jedes Mal, wenn sie aufgerufen wird, einen Wert zwischen `0` (einschließlich) und `1` (ausschließlich) zurück. Für eine 50/50-Chance verwenden Sie `Math.random() >= .5` in Ihrem Ausdruck. Statistisch gesehen gibt dieser Ausdruck 50 % der Zeit `true` und die anderen 50 % `false` zurück. Ersetzen Sie in der Render-Methode `null` durch den oben genannten Ausdruck, um die Variablendeklaration abzuschließen.

Nun haben Sie einen Ausdruck, mit dem Sie eine zufällige Entscheidung im Code treffen können. Als Nächstes müssen Sie dies implementieren. Rendern Sie die `Results`-Komponente als Kind von `GameOfChance` und übergeben Sie `expression` als Prop mit dem Namen `fiftyFifty`. Schreiben Sie in der `Results`-Komponente einen ternären Ausdruck, um das `h1`-Element mit dem Text `You Win!` oder `You Lose!` basierend auf der `fiftyFifty`-Prop zu rendern, die von `GameOfChance` übergeben wird. Stellen Sie schließlich sicher, dass die `handleClick()`-Methode jeden Zug korrekt zählt, damit der Benutzer weiß, wie oft er gespielt hat. Dies dient auch dazu, dem Benutzer zu zeigen, dass die Komponente tatsächlich aktualisiert wurde, falls er zweimal hintereinander gewinnt oder verliert.

# --hints--

Die `GameOfChance`-Komponente sollte existieren und auf der Seite gerendert werden.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance` sollte ein einzelnes `button`-Element zurückgeben.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance` sollte eine einzelne Instanz der `Results`-Komponente zurückgeben, die eine Prop namens `fiftyFifty` hat.

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

Der Zustand von `GameOfChance` sollte mit einer Eigenschaft `counter` initialisiert werden, die auf den Wert `1` gesetzt ist.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

Wenn die `GameOfChance`-Komponente zum ersten Mal im DOM gerendert wird, sollte ein `p`-Element mit dem inneren Text `Turn: 1` zurückgegeben werden.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

Jedes Mal, wenn die Taste gedrückt wird, sollte der Zählerzustand um 1 erhöht werden, und ein einzelnes `p`-Element sollte im DOM gerendert werden, das den Text `Turn: N` enthält, wobei `N` der Wert des Zählerzustands ist.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

Wenn die `GameOfChance`-Komponente zum ersten Mal im DOM gemountet wird und jedes Mal, wenn die Taste danach gedrückt wird, sollte ein einzelnes `h1`-Element zurückgegeben werden, das zufällig entweder `You Win!` oder `You Lose!` rendert. Hinweis: Dies kann zufällig fehlschlagen. Falls das passiert, versuchen Sie es bitte erneut.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
