---
id: 5a24c314108439a4d403617a
title: Zustand als Props an Kindkomponenten übergeben
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

Sie haben in vorherigen Aufgaben viele Beispiele gesehen, in denen Props an untergeordnete JSX-Elemente und React-Komponenten übergeben wurden. Vielleicht fragen Sie sich, woher diese Props kommen. Ein gängiges Muster ist, eine zustandsbehaftete Komponente zu haben, die den `state` enthält, der für Ihre Anwendung wichtig ist, und die dann Kindkomponenten rendert. Sie möchten, dass diese Komponenten Zugriff auf einige Teile dieses `state` haben, die als Props übergeben werden.

Zum Beispiel haben Sie vielleicht eine `App`-Komponente, die eine `Navbar` rendert, neben anderen Komponenten. In Ihrer `App` haben Sie `state`, die viele Benutzerinformationen enthält, aber die `Navbar` benötigt nur Zugriff auf den Benutzernamen, um ihn anzuzeigen. Dieses Stück `state` übergeben Sie als Prop an die `Navbar`-Komponente.

Dieses Muster veranschaulicht einige wichtige Paradigmen in React. Das erste ist der *unidirektionale Datenfluss*. Zustand fließt in eine Richtung den Baum Ihrer Anwendungs-Komponenten hinunter, von der zustandsbehafteten Elternkomponente zu den Kindkomponenten. Die Kindkomponenten erhalten nur die Zustandsdaten, die sie benötigen. Das zweite ist, dass komplexe zustandsbehaftete Anwendungen in nur wenige oder vielleicht eine einzige zustandsbehaftete Komponente aufgeteilt werden können. Der Rest Ihrer Komponenten erhält einfach Zustand vom Elternteil als Props und rendert daraus eine Benutzeroberfläche. So entsteht eine Trennung, bei der die Zustandsverwaltung in einem Teil des Codes und das UI-Rendering in einem anderen Teil gehandhabt wird. Dieses Prinzip, Zustandslogik von UI-Logik zu trennen, ist einer der Schlüsselgrundsätze von React. Wenn es richtig angewendet wird, erleichtert es das Design komplexer, zustandsbehafteter Anwendungen erheblich.

# --instructions--

Die `MyApp`-Komponente ist zustandsbehaftet und rendert eine `Navbar`-Komponente als Kind. Übergeben Sie die `name`-Eigenschaft in ihrem `state` an die Kindkomponente und zeigen Sie dann die `name` im `h1`-Tag an, der Teil der `Navbar`-Render-Methode ist. `name` sollte nach dem Text `Hello, my name is:` erscheinen.

# --hints--

Die `MyApp`-Komponente sollte mit einer `Navbar`-Komponente darin rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

Die `Navbar`-Komponente sollte die `MyApp`-Zustandseigenschaft `name` als Props erhalten.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

Das `h1`-Element in `Navbar` sollte die `name`-Prop rendern.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
