---
id: 5a24c314108439a4d403617b
title: Callback als Props übergeben
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

Sie können `state` als Props an Kind-Komponenten übergeben, aber Sie sind nicht darauf beschränkt, nur Daten zu übergeben. Sie können auch Handler-Funktionen oder jede Methode, die in einer React-Komponente definiert ist, an eine Kind-Komponente übergeben. So ermöglichen Sie es Kind-Komponenten, mit ihren Eltern-Komponenten zu interagieren. Sie übergeben Methoden an ein Kind genauso wie eine normale Prop. Sie wird mit einem Namen zugewiesen, und Sie haben Zugriff auf diesen Methodennamen unter `this.props` in der Kind-Komponente.

# --instructions--

Im Code-Editor sind drei Komponenten skizziert. Die `MyApp`-Komponente ist die Elternkomponente, die die Kind-Komponenten `GetInput` und `RenderInput` rendern wird. Fügen Sie die `GetInput`-Komponente zur render-Methode in `MyApp` hinzu, und übergeben Sie ihr eine Prop namens `input`, die auf `inputValue` aus `MyApp`s `state` zugewiesen ist. Erstellen Sie außerdem eine Prop namens `handleChange` und übergeben Sie ihr den Eingabe-Handler `handleChange`.

Fügen Sie anschließend `RenderInput` zur render-Methode in `MyApp` hinzu, erstellen Sie eine Prop namens `input` und übergeben Sie ihr `inputValue` aus `state`. Sobald Sie fertig sind, können Sie im Eingabefeld `input` in der `GetInput`-Komponente tippen, was dann die Handler-Methode im Elternteil über Props aufruft. Dadurch wird die Eingabe im `state` des Elternteils aktualisiert, die als Props an beide Kinder übergeben wird. Beobachten Sie, wie die Daten zwischen den Komponenten fließen und wie die einzige Quelle der Wahrheit `state` der Elternkomponente bleibt. Zugegeben, dieses Beispiel ist etwas konstruiert, sollte aber veranschaulichen, wie Daten und Callbacks zwischen React-Komponenten übergeben werden können.

# --hints--

Die `MyApp`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

Die `GetInput`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

Die `RenderInput`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

Die `GetInput`-Komponente sollte die `MyApp`-Status-Eigenschaft `inputValue` als Props erhalten und ein `input`-Element enthalten, das den `MyApp`-Status ändert.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

Die `RenderInput`-Komponente sollte die `MyApp`-Status-Eigenschaft `inputValue` als Props erhalten.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
