---
id: 5a24c314108439a4d4036180
title: Optimieren von Neurenderings mit shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

Bisher rendert sich jede Komponente neu und alle ihre Kinder, wenn sie neue `state` oder neue `props` erhält. Das ist normalerweise in Ordnung. React stellt jedoch eine Lifecycle-Methode bereit, die Sie aufrufen können, wenn Kindkomponenten neue `state` oder `props` erhalten, und mit der Sie speziell angeben können, ob die Komponenten aktualisiert werden sollen oder nicht. Die Methode heißt `shouldComponentUpdate()` und nimmt `nextProps` und `nextState` als Parameter entgegen.

Diese Methode ist eine nützliche Möglichkeit, die Leistung zu optimieren. Zum Beispiel ist das Standardverhalten, dass Ihre Komponente neu rendert, wenn sie neue `props` erhält, selbst wenn sich die `props` nicht geändert haben. Sie können `shouldComponentUpdate()` verwenden, um dies zu verhindern, indem Sie die `props` vergleichen. Die Methode muss einen `boolean` Wert zurückgeben, der React mitteilt, ob die Komponente aktualisiert werden soll oder nicht. Sie können die aktuellen Props (`this.props`) mit den nächsten Props (`nextProps`) vergleichen, um zu bestimmen, ob Sie aktualisieren müssen oder nicht, und entsprechend `true` oder `false` zurückgeben.

# --instructions--

Die `shouldComponentUpdate()` Methode ist in einer Komponente namens `OnlyEvens` hinzugefügt. Derzeit gibt diese Methode `true` zurück, sodass `OnlyEvens` jedes Mal neu rendert, wenn sie neue `props` erhält. Ändern Sie die Methode so, dass `OnlyEvens` nur aktualisiert wird, wenn der `value` ihrer neuen Props gerade ist. Klicken Sie auf die `Add` Schaltfläche und beobachten Sie die Reihenfolge der Ereignisse in der Konsole Ihres Browsers, während die Lifecycle-Hooks ausgelöst werden.

# --hints--

Die `Controller` Komponente sollte die `OnlyEvens` Komponente als Kind rendern.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

Die `shouldComponentUpdate` Methode sollte in der `OnlyEvens` Komponente definiert sein.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

Die `OnlyEvens` Komponente sollte ein `h1` Tag zurückgeben, das den Wert von `this.props.value` rendert.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` sollte nur neu rendern, wenn `nextProps.value` gerade ist.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
