---
id: 5a24c314108439a4d4036183
title: Fortgeschrittenes JavaScript in der React-Render-Methode verwenden
challengeType: 6
forumTopicId: 301415
dashedName: use-advanced-javascript-in-react-render-method
---

# --description--

In vorherigen Aufgaben haben Sie gelernt, wie Sie JavaScript-Code mithilfe von geschweiften Klammern, `{ }`, in JSX-Code einfügen können, um Aufgaben wie das Zugreifen auf Props, das Übergeben von Props, das Zugreifen auf den State, das Einfügen von Kommentaren in Ihren Code und zuletzt das Stylen Ihrer Komponenten zu erledigen. Dies sind alles gängige Anwendungsfälle, um JavaScript in JSX zu verwenden, aber sie sind nicht die einzigen Möglichkeiten, wie Sie JavaScript-Code in Ihren React-Komponenten nutzen können.

Sie können auch JavaScript direkt in Ihren `render`-Methoden schreiben, vor der `return`-Anweisung, ***ohne*** es in geschweifte Klammern einzufügen. Das liegt daran, dass es sich noch nicht innerhalb des JSX-Codes befindet. Wenn Sie später eine Variable im JSX-Code *innerhalb* der `return`-Anweisung verwenden möchten, setzen Sie den Variablennamen in geschweifte Klammern.

# --instructions--

Im bereitgestellten Code enthält die `render`-Methode ein Array mit 20 Phrasen, die die Antworten des klassischen Magic Eight Ball-Spielzeugs aus den 1980er Jahren darstellen. Das Klick-Ereignis des Buttons ist an die `ask`-Methode gebunden, sodass bei jedem Klick auf den Button eine Zufallszahl generiert und als `randomIndex` im State gespeichert wird. Löschen Sie in Zeile 52 den String `change me!` und weisen Sie die `answer`-Konstante neu zu, damit Ihr Code bei jeder Aktualisierung der Komponente zufällig auf einen anderen Index des `possibleAnswers`-Arrays zugreift. Fügen Sie schließlich die `answer`-Konstante innerhalb der `p`-Tags ein.

# --hints--

Die `MagicEightBall`-Komponente sollte existieren und auf der Seite gerendert werden.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall)).find('MagicEightBall')
    .length,
  1
);
```

Das erste Kind von `MagicEightBall` sollte ein `input`-Element sein.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall))
    .children()
    .childAt(0)
    .name(),
  'input'
);
```

Das dritte Kind von `MagicEightBall` sollte ein `button`-Element sein.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MagicEightBall))
    .children()
    .childAt(2)
    .name(),
  'button'
);
```

Der State von `MagicEightBall` sollte mit einer Eigenschaft `userInput` und einer Eigenschaft `randomIndex` initialisiert werden, die beide den Wert eines leeren Strings haben.

```js
assert(
  Enzyme.mount(React.createElement(MagicEightBall)).state('randomIndex') ===
    '' &&
    Enzyme.mount(React.createElement(MagicEightBall)).state('userInput') === ''
);
```

Wenn `MagicEightBall` zum ersten Mal im DOM gemountet wird, sollte es ein leeres `p`-Element zurückgeben.

```js
assert(
  Enzyme.mount(React.createElement(MagicEightBall)).find('p').length === 1 &&
    Enzyme.mount(React.createElement(MagicEightBall)).find('p').text() === ''
);
```

Wenn Text in das `input`-Element eingegeben wird und der Button geklickt wird, sollte die `MagicEightBall`-Komponente ein `p`-Element zurückgeben, das ein zufälliges Element aus dem `possibleAnswers`-Array enthält.

```js
(() => {
  const comp = Enzyme.mount(React.createElement(MagicEightBall));
  const simulate = () => {
    comp.find('input').simulate('change', { target: { value: 'test?' } });
    comp.find('button').simulate('click');
  };
  const result = () => comp.find('p').text();
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const actualAnswers = [
    _1_val,
    _2_val,
    _3_val,
    _4_val,
    _5_val,
    _6_val,
    _7_val,
    _8_val,
    _9_val,
    _10_val
  ];
  const hasIndex = actualAnswers.filter(
    (answer, i) => possibleAnswers.indexOf(answer) !== -1
  );
  const notAllEqual = new Set(actualAnswers);
  assert(notAllEqual.size > 1 && hasIndex.length === 10);
})();
```

# --before-all--

```js
var possibleAnswers = [
  'It is certain',
  'It is decidedly so',
  'Without a doubt',
  'Yes, definitely',
  'You may rely on it',
  'As I see it, yes',
  'Outlook good',
  'Yes',
  'Signs point to yes',
  'Reply hazy try again',
  'Ask again later',
  'Better not tell you now',
  'Cannot predict now',
  'Concentrate and ask again',
  "Don't count on it",
  'My reply is no',
  'My sources say no',
  'Outlook not so good',
  'Very doubtful',
  'Most likely'
];
```

# --seed--

## --seed-contents--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class MagicEightBall extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userInput: '',
      randomIndex: ''
    };
    this.ask = this.ask.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  ask() {
    if (this.state.userInput) {
      this.setState({
        randomIndex: Math.floor(Math.random() * 20),
        userInput: ''
      });
    }
  }
  handleChange(event) {
    this.setState({
      userInput: event.target.value
    });
  }
  render() {
    const possibleAnswers = [
      'It is certain',
      'It is decidedly so',
      'Without a doubt',
      'Yes, definitely',
      'You may rely on it',
      'As I see it, yes',
      'Outlook good',
      'Yes',
      'Signs point to yes',
      'Reply hazy try again',
      'Ask again later',
      'Better not tell you now',
      'Cannot predict now',
      'Concentrate and ask again',
      "Don't count on it",
      'My reply is no',
      'My sources say no',
      'Most likely',
      'Outlook not so good',
      'Very doubtful'
    ];
    const answer = 'change me!'; // Change this line
    return (
      <div>
        <input
          type='text'
          value={this.state.userInput}
          onChange={this.handleChange}
          style={inputStyle}
        />
        <br />
        <button onClick={this.ask}>Ask the Magic Eight Ball!</button>
        <br />
        <h3>Answer:</h3>
        <p>
          {/* Change code below this line */}

          {/* Change code above this line */}
        </p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
const inputStyle = {
  width: 235,
  margin: 5
};

class MagicEightBall extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      userInput: '',
      randomIndex: ''
    };
    this.ask = this.ask.bind(this);
    this.handleChange = this.handleChange.bind(this);
  }
  ask() {
    if (this.state.userInput) {
      this.setState({
        randomIndex: Math.floor(Math.random() * 20),
        userInput: ''
      });
    }
  }
  handleChange(event) {
    this.setState({
      userInput: event.target.value
    });
  }
  render() {
    const possibleAnswers = [
      'It is certain',
      'It is decidedly so',
      'Without a doubt',
      'Yes, definitely',
      'You may rely on it',
      'As I see it, yes',
      'Outlook good',
      'Yes',
      'Signs point to yes',
      'Reply hazy try again',
      'Ask again later',
      'Better not tell you now',
      'Cannot predict now',
      'Concentrate and ask again',
      "Don't count on it",
      'My reply is no',
      'My sources say no',
      'Outlook not so good',
      'Very doubtful',
      'Most likely'
    ];
    const answer = possibleAnswers[this.state.randomIndex];
    return (
      <div>
        <input
          type='text'
          value={this.state.userInput}
          onChange={this.handleChange}
          style={inputStyle}
        />
        <br />
        <button onClick={this.ask}>Ask the Magic Eight Ball!</button>
        <br />
        <h3>Answer:</h3>
        <p>{answer}</p>
      </div>
    );
  }
}
```
