---
id: 5a24c314108439a4d4036184
title: Rendern mit einer If-Else-Bedingung
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

Eine weitere Möglichkeit, JavaScript zu verwenden, um Ihre gerenderte Ansicht zu steuern, besteht darin, die gerenderten Elemente an eine Bedingung zu knüpfen. Wenn die Bedingung wahr ist, wird eine Ansicht gerendert. Wenn sie falsch ist, wird eine andere Ansicht gerendert. Dies können Sie mit einer normalen `if/else`-Anweisung in der `render()`-Methode einer React-Komponente tun.

# --instructions--

MyComponent enthält einen `boolean` in seinem State, der verfolgt, ob Sie ein Element in der Benutzeroberfläche anzeigen möchten oder nicht. Der `button` schaltet den Zustand dieses Werts um. Derzeit rendert es jedes Mal dieselbe Benutzeroberfläche. Schreiben Sie die `render()`-Methode mit einer `if/else`-Anweisung so um, dass, wenn `display` `true` ist, Sie die aktuelle Auszeichnung zurückgeben. Andernfalls geben Sie die Auszeichnung ohne das `h1`-Element zurück.

**Hinweis:** Sie müssen eine `if/else` schreiben, um die Tests zu bestehen. Die Verwendung des ternären Operators wird hier nicht akzeptiert.

# --hints--

`MyComponent` sollte existieren und rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

Wenn `display` auf `true` gesetzt ist, sollten ein `div`, `button` und `h1` gerendert werden.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

Wenn `display` auf `false` gesetzt ist, sollten nur ein `div` und `button` gerendert werden.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

Die render-Methode sollte eine `if/else`-Anweisung verwenden, um die Bedingung von `this.state.display` zu prüfen.

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
