---
id: 5a24c314108439a4d4036176
title: "Zustand verwenden, um ein Element umzuschalten"
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

Manchmal müssen Sie den vorherigen Zustand kennen, wenn Sie den Zustand aktualisieren. Allerdings können Zustandsaktualisierungen asynchron sein – das bedeutet, React kann mehrere `setState()`-Aufrufe zu einer einzigen Aktualisierung zusammenfassen. Das heißt, Sie können sich beim Berechnen des nächsten Werts nicht auf den vorherigen Wert von `this.state` oder `this.props` verlassen. Deshalb sollten Sie keinen Code wie diesen verwenden:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

Stattdessen sollten Sie `setState` eine Funktion übergeben, die Ihnen Zugriff auf Zustand und Props erlaubt. Die Verwendung einer Funktion mit `setState` garantiert, dass Sie mit den aktuellsten Werten von Zustand und Props arbeiten. Das bedeutet, dass der obige Code so umgeschrieben werden sollte:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

Sie können auch eine Form ohne `props` verwenden, wenn Sie nur den `state` benötigen:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

Beachten Sie, dass Sie das Objektliteral in Klammern setzen müssen, sonst interpretiert JavaScript es als Codeblock.

# --instructions--

`MyComponent` hat eine `visibility`-Eigenschaft, die mit `false` initialisiert wird. Die Render-Methode gibt eine Ansicht zurück, wenn der Wert von `visibility` wahr ist, und eine andere Ansicht, wenn er falsch ist.

Derzeit gibt es keine Möglichkeit, die `visibility`-Eigenschaft im `state` der Komponente zu aktualisieren. Der Wert soll zwischen true und false hin- und herschalten. Es gibt einen Klick-Handler auf dem Button, der eine Klassenmethode namens `toggleVisibility()` auslöst. Übergeben Sie eine Funktion an `setState`, um diese Methode so zu definieren, dass die `state` von `visibility` beim Aufrufen der Methode auf den gegenteiligen Wert umschaltet. Wenn `visibility` `false` ist, setzt die Methode ihn auf `true` und umgekehrt.

Klicken Sie abschließend auf den Button, um das bedingte Rendern der Komponente basierend auf ihrem `state` zu sehen.

**Hinweis:** Vergessen Sie nicht, das Schlüsselwort `this` in der `constructor` an die Methode zu binden!

# --hints--

`MyComponent` sollte ein `div`-Element zurückgeben, das ein `button` enthält.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

Der Zustand von `MyComponent` sollte mit einer `visibility`-Eigenschaft initialisiert werden, die auf `false` gesetzt ist.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

Das Klicken auf das Button-Element sollte die `visibility`-Eigenschaft im Zustand zwischen `true` und `false` umschalten und das `h1`-Element bedingt rendern.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

Eine anonyme Funktion sollte an `setState` übergeben werden.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this` sollte nicht innerhalb von `setState` verwendet werden.

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
