---
id: 5a24c314108439a4d4036172
title: Zustand in der Benutzeroberfläche auf eine andere Weise rendern
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

Es gibt eine andere Möglichkeit, in einer Komponente auf `state` zuzugreifen. In der `render()`-Methode können Sie vor der `return`-Anweisung direkt JavaScript schreiben. Zum Beispiel könnten Sie Funktionen deklarieren, auf Daten aus `state` oder `props` zugreifen, Berechnungen mit diesen Daten durchführen und so weiter. Anschließend können Sie beliebige Daten Variablen zuweisen, auf die Sie in der `return`-Anweisung zugreifen können.

# --instructions--

Definieren Sie in der `MyComponent`-Render-Methode eine `const` namens `name` und setzen Sie sie gleich dem Namen-Wert im `state` der Komponente. Da Sie in diesem Teil des Codes direkt JavaScript schreiben können, müssen Sie diese Referenz nicht in geschweifte Klammern einschließen.

Rendern Sie anschließend in der return-Anweisung diesen Wert in einem `h1`-Tag unter Verwendung der Variablen `name`. Denken Sie daran, dass Sie in der return-Anweisung die JSX-Syntax verwenden müssen (geschweifte Klammern für JavaScript).

# --hints--

`MyComponent` sollte einen Schlüssel `name` mit dem Wert `freeCodeCamp` haben, der im Zustand gespeichert ist.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` sollte ein `h1`-Überschriftselement rendern, das in einem einzigen `div` eingeschlossen ist.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

Das gerenderte `h1`-Tag sollte eine Referenz auf `{name}` enthalten.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

Das gerenderte `h1`-Überschriftselement sollte Text enthalten, der aus dem Zustand der Komponente gerendert wird.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
