---
id: 5a24c314108439a4d4036160
title: Eine HTML-Klasse in JSX definieren
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

Jetzt, wo Sie sich mit dem Schreiben von JSX vertraut machen, fragen Sie sich vielleicht, worin der Unterschied zu HTML besteht.

Bisher mag es so erscheinen, als seien HTML und JSX genau gleich.

Ein wesentlicher Unterschied in JSX ist, dass Sie das Wort `class` nicht mehr verwenden können, um HTML-Klassen zu definieren. Das liegt daran, dass `class` ein reserviertes Schlüsselwort in JavaScript ist. Stattdessen verwendet JSX `className`.

Tatsächlich wird die Benennungskonvention für alle HTML-Attribute und Ereignis-Referenzen in JSX zu CamelCase. Zum Beispiel heißt ein Klick-Ereignis in JSX `onClick` statt `onclick`. Ebenso wird `onchange` zu `onChange`. Obwohl dies ein subtiler Unterschied ist, ist es wichtig, ihn für die weitere Arbeit im Hinterkopf zu behalten.

# --instructions--

Weisen Sie dem in dem JSX-Code bereitgestellten `myDiv` eine Klasse mit dem Namen `div` zu.

# --hints--

Die Konstante `JSX` sollte ein `div`-Element zurückgeben.

```js
assert.strictEqual(JSX.type, 'div');
```

Das `div` sollte eine Klasse mit dem Namen `myDiv` haben.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
