---
id: 6a89c69c99298438d6fbe60b
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Agora adicione o caso recursivo. Depois da declaração `if`, crie uma variável chamada `count_list`. Atribua a ela o resultado da chamada de `countup` com `number - 1`.

Cada chamada usa um número menor, então a recursão avança em direção ao caso base.

# --hints--

Você deve atribuir `countup(number - 1)` a uma variável chamada `count_list`.

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
