---
id: 6a89c69c6af682856558d78b
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

A chamada recursiva retorna uma lista contendo os números menores. Use `append()` para adicionar `number` a `count_list`, depois retorne `count_list`.

# --hints--

Você deve adicionar `number` a `count_list`.

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

Você deve retornar `count_list` da sua função `countup`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

Sua função `countup` deve retornar `[1, 2, 3]` quando chamada com `3`.

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
