---
id: 6a89c69c111694edc6d1c628
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Teste sua função `countup` passando `5` para ela dentro de uma chamada `print()`. A saída deve ser `[1, 2, 3, 4, 5]`.

Com isso, a oficina da função countup está completa.

# --hints--

Você deve fazer a chamada `countup(5)` dentro da função `print()`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
