---
id: 6a89c69c8dbfda74151fa0b6
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nesta oficina, você vai criar uma função `countup` que usa recursão para criar uma lista de números de `1` até um número dado.

Comece definindo uma função chamada `countup` com um parâmetro chamado `number`. Adicione a palavra-chave `pass` ao corpo da função para que ela seja válida enquanto estiver vazia.

# --hints--

Você deve ter uma função chamada `countup` com um parâmetro chamado `number`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

Sua função `countup` deve conter a palavra-chave `pass`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
