---
id: 6a89c69c91bc7365fbd4b733
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Adicione o caso base à sua função `countup`. Substitua `pass` por uma declaração `if` que verifica se `number` é menor que `1`. Dentro da declaração `if`, retorne uma lista vazia.

# --hints--

Sua função `countup` deve retornar uma lista vazia quando `number` for menor que `1`.

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

Sua função `countup` não deve retornar uma lista vazia quando `number` for maior ou igual a `1`.

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
