---
id: 69a5f35669099ed52f8563b1
title: Ein intelligentes Vorratskammer-Nachfüllprogramm erstellen
challengeType: 26
dashedName: lab-smart-pantry-restocker
---

# --description--

In diesem Labor erstellen Sie ein kleines Programm zur Verwaltung der Vorratskammer unter Verwendung grundlegender JavaScript-Konzepte wie Arrays, Objekte, Schleifen und Bedingungen.

Sie simulieren den Erhalt einer Lieferung von Vorratsartikeln, entscheiden, was mit jedem Artikel zu tun ist, und organisieren die Ergebnisse zur Lagerung.

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

Das `rawData`-Array enthält durch Pipes getrennte Strings im Format `sku|name|qty|expires|zone`, wobei `zone` optional ist.

**User Stories:**

1. Sie sollten eine `parseShipment(rawData)`-Funktion implementieren, die ein Array von Strings entgegennimmt und ein Array von Objekten mit `{ sku, name, qty, expires, zone }`-Eigenschaften zurückgibt.

   - Doppelte `sku`-Werte in der Lieferung sollten ignoriert werden.
   - Wenn der `zone`-Abschnitt nicht angegeben ist, sollte er standardmäßig auf `"general"` gesetzt werden.
   - Der Wert von `qty` sollte in eine Zahl umgewandelt werden.

2. Sie sollten eine `planRestock(pantry, shipment)`-Funktion implementieren, die die aktuelle Vorratskammer mit der eingehenden Lieferung vergleicht und ein Array von Aktionen in der Form `{ type, item }` zurückgibt, wobei `type` einer von `"restock"`, `"discard"` oder `"donate"` ist und `item` das geparste Lieferobjekt ist.

   Der Parameter `pantry` ist ein Array von Objekten mit derselben Struktur wie ein geparstes Lieferobjekt (`{ sku, name, qty, expires, zone }`).

- Wenn ein Lieferungsartikel eine `qty` von `0` oder weniger hat, sollte der Aktion `type` `"discard"` sein, unabhängig davon, ob der Artikel in der Vorratskammer vorhanden ist.
   - Andernfalls, wenn die `sku` des Lieferungsartikels bereits in der Vorratskammer existiert, sollte der Aktion `type` `"restock"` sein.
   - Andernfalls (wenn die `sku` des Lieferungsartikels nicht in der Vorratskammer existiert), sollte der Aktion `type` `"donate"` sein.

3. Sie sollten eine `groupByZone(actions)`-Funktion implementieren, die die Aktionen basierend auf der `zone`-Eigenschaft jedes Artikels in Lagerzonen gruppiert. Die Funktion sollte ein Objekt zurückgeben, bei dem jeder Schlüssel ein Zonenname ist und der Wert ein Array von Aktionen, die zu dieser Zone gehören. Wenn Aktionen beispielsweise Artikel mit den Zonen `"fridge"` und `"pantry"` enthalten, sollte das Ergebnis `{ fridge: [...], pantry: [...] }` sein.

4. Sie sollten eine `clonePantry(pantry)`-Funktion implementieren, die eine tiefe Kopie der Vorratskammer zurückgibt, sodass Planungsänderungen die ursprüngliche Liste nicht beeinflussen. Eine tiefe Kopie bedeutet, ein neues Array mit neuen Objekten zu erstellen, sodass Änderungen an der Kopie die ursprüngliche Vorratskammer nicht verändern.

5. Sie sollten alle Funktionen zusammen verwenden, um eine Lieferung zu verarbeiten und das endgültige gruppierte Ergebnisobjekt in der Konsole auszugeben.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Funktion mit dem Namen `parseShipment` definieren, die ein Array von Strings namens `rawData` als Parameter akzeptiert.

``` js
assert.isFunction(parseShipment);
assert.lengthOf(parseShipment, 1);
```

Ihre `parseShipment`-Funktion sollte die Lieferungs-Strings im Array in Objekte mit den Eigenschaften `sku`, `name`, `qty`, `expires` und `zone` umwandeln.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);

assert.isArray(result);
assert.isObject(result[0]);
assert.deepEqual(result[0], {
  sku: "A10",
  name: "Tomatoes",
  qty: 5,
  expires: "2027-01-01",
  zone: "fridge"
});
```

Doppelte SKUs in der Lieferung sollten ignoriert werden.

``` js
const rawData = [
  "C32|Eggs|3|2027-01-01",
  "C32|Eggs|3|2027-01-01"
];

const result = parseShipment(rawData);

assert.lengthOf(result, 1);
assert.strictEqual(result[0].sku, "C32");
```

Wenn `zone` in einem Lieferungs-String nicht vorhanden ist, sollten Sie es auf `general` setzen.

``` js
const rawData = ["B21|Bananas|10|2027-01-01"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].zone, "general");
```

Der Wert von `qty` sollte in eine Zahl umgewandelt werden.

``` js
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];
const result = parseShipment(rawData);
assert.strictEqual(result[0].qty, 5);
assert.isNumber(result[0].qty);
```

Sie sollten eine Funktion mit dem Namen `planRestock` definieren, die zwei Parameter akzeptiert: `pantry` und `shipment`.

``` js
assert.isFunction(planRestock);
assert.lengthOf(planRestock, 2);
```

Ihre `planRestock`-Funktion sollte ein Array von Aktionen mit den Eigenschaften `type` und `item` zurückgeben.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 6, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.isArray(result);
assert.isObject(result[0]);
assert.property(result[0], "type");
assert.property(result[0], "item");
```

Artikel mit einer Menge von `0` oder weniger in der Lieferung sollten eine Aktion mit dem Typ `discard` erzeugen, unabhängig davon, ob der Artikel in der Vorratskammer vorhanden ist.

``` js
const pantry = [{ sku: "D43", name: "Pineapples", qty: 2, expires: "2027-01-01", zone: "general" }];
const shipment = [{ sku: "D43", name: "Pineapples", qty: 0, expires: "2027-01-01", zone: "general" },
                  { sku: "E54", name: "Peppers", qty: -1, expires: "2027-01-01", zone: "fridge" }
                ];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "discard");
assert.strictEqual(result[1].type, "discard");
```

Wenn ein Lieferungsartikel bereits in der Vorratskammer existiert (gleiche `sku`), sollte die Aktion `restock` sein.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const shipment = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "restock");
```

Wenn ein Lieferungsartikel nicht in der Vorratskammer existiert, sollte die Aktion `"donate"` sein.

``` js
const pantry = [];
const shipment = [{ sku: "C32", name: "Eggs", qty: 3, expires: "2027-01-01", zone: "fridge" }];

const result = planRestock(pantry, shipment);

assert.strictEqual(result[0].type, "donate");
```

Sie sollten eine Funktion mit dem Namen `groupByZone` definieren, die einen Parameter namens `actions` akzeptiert.

``` js
assert.isFunction(groupByZone);
assert.lengthOf(groupByZone, 1);
```

Ihre Funktion `groupByZone` sollte die Aktionen gruppiert nach der `zone`-Eigenschaft jedes Objekts zurückgeben.

``` js

const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "restock", item: { sku: "B1", zone: "pantry" } }, 
];

const result = groupByZone(actions);

assert.isObject(result);
assert.isArray(result.fridge);
assert.isArray(result.pantry);
```

Ihre `groupByZone`-Funktion sollte Aktionen mit dem richtigen Inhalt und der richtigen Anzahl korrekt gruppieren.

``` js
const actions = [
  { type: "restock", item: { sku: "A1", zone: "fridge" } },
  { type: "donate",  item: { sku: "B1", zone: "fridge" } },
  { type: "discard", item: { sku: "C1", zone: "pantry" } },
];

const result = groupByZone(actions);

assert.lengthOf(result.fridge, 2);
assert.lengthOf(result.pantry, 1);
assert.strictEqual(result.fridge[0].item.sku, "A1");
assert.strictEqual(result.pantry[0].type, "discard");
```

Sie sollten eine Funktion mit dem Namen `clonePantry` definieren, die einen Parameter namens `pantry` akzeptiert.

``` js
assert.isFunction(clonePantry);
assert.lengthOf(clonePantry, 1);
```

`clonePantry` sollte ein neues Array zurückgeben, anstatt das ursprüngliche Vorratskammer-Array.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.isArray(pantryClone);
assert.notStrictEqual(pantryClone, pantry);
```

Die Objekte innerhalb der geklonten Vorratskammer sollten ebenfalls neue Objekte sein (tiefe Kopie).

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const pantryClone = clonePantry(pantry);

assert.notStrictEqual(pantryClone[0], pantry[0]);
assert.deepEqual(pantryClone[0], pantry[0]);

pantryClone[0].qty = 100;
assert.strictEqual(pantry[0].qty, 4);
```

Die Funktionen sollten zusammenarbeiten, um eine Lieferung zu verarbeiten und die resultierenden Aktionen zu gruppieren.

``` js
const pantry = [{ sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" }];
const rawData = ["A10|Tomatoes|5|2027-01-01|fridge"];

const shipment = parseShipment(rawData);
const pantryClone = clonePantry(pantry);
const actions = planRestock(pantryClone, shipment);
const grouped = groupByZone(actions);

assert.isObject(grouped);
assert.strictEqual(grouped.fridge[0].item.qty, 5);
```

Sie sollten die resultierenden Aktionen gruppiert nach Zonen in der Konsole ausgeben.

``` js
const groupedLog = getLogs().find(log => {
  if (!log || typeof log !== "object" || Array.isArray(log)) return false;

  return Object.values(log).some(zoneActions => {
    if (!Array.isArray(zoneActions)) return false;

    return zoneActions.some(action =>
      action &&
      ["restock", "discard", "donate"].includes(action.type) &&
      action.item &&
      typeof action.item.sku === "string"
    );
  });
});

assert.isOk(groupedLog, "Log the grouped actions object to the console.");
```

# --seed--

## --seed-contents--

```js
const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01",
  "B21|Bananas|10|2027-01-01",
  "C32|Eggs|3|2027-01-01|fridge",
  "C32|Eggs|3|2027-01-01",
  "D43|Pineapples|0|2027-01-01",
  "E54|Peppers|-1|2027-01-01|fridge"
];
```

# --solutions--

```js

function parseShipment(rawData) {

  // Early return if rawData is not an Array
  if (!Array.isArray(rawData)) return [];

  const result = [];
  const seenSKUs = [];

  for (let line of rawData) {

    const parts = line.split("|");

    const sku = parts[0].trim();

    if (seenSKUs.includes(sku)) continue; // Avoid duplicates

    const name = parts[1].trim();
    const qty = Number(parts[2].trim()); // Parse qty into a number
    const expires = parts[3].trim();
    const zone = ( parts[4] ) ? parts[4].trim() : "general";

    result.push({
      sku: sku,
      name: name,
      qty: qty,
      expires: expires,
      zone: zone
    });

    seenSKUs.push(sku);

  }

  return result;
}

function planRestock(pantry, shipment) {
  const actions = [];

  for (let item of shipment) {

    var exists = false;

    for (let stock of pantry)
    {
      if (stock.sku === item.sku)
      {
        exists = true;
      }
    }

    if (item.qty > 0) {
      if (exists) {
        actions.push({ type: "restock", item: item });
      } else{
        actions.push({ type: "donate", item: item });
      }
    } else {
      actions.push({ type: "discard", item: item });
    }
  }

  return actions;
}

function groupByZone(actions) {

  const byZone = {};

  for (const action of actions) {

    const zone = action.item.zone;

    if (!byZone[zone]) {
      byZone[zone] = [];
    }

    byZone[zone].push(action);
  }

  return byZone;

}

function clonePantry(pantry) {
  const pantryClone = [];

  for (let item of pantry) {
    pantryClone.push({
      sku: item.sku,
      name: item.name,
      qty: item.qty,
      expires: item.expires,
      zone: item.zone
    });
  }

  return pantryClone;
}

const pantry = [
  { sku: "A10", name: "Tomatoes", qty: 4, expires: "2027-01-01", zone: "fridge" },
  { sku: "D43", name: "Pineapples", qty: 2, expires: "2020-01-01", zone: "general" }
];

const rawData = [
  "A10|Tomatoes|5|2027-01-01", // Restock existing item
  "B21|Bananas|10|2027-01-01", // Donate new item without zone
  "C32|Eggs|3|2027-01-01|fridge", // Donate to a defined zone
  "C32|Eggs|3|2027-01-01", // Duplicated SKU in shipment
  "D43|Pineapples|0|2027-01-01", // Discard with quantity of 0
  "E54|Peppers|-1|2027-01-01|fridge" // Discard even if it's not in pantry
];

const shipment = parseShipment(rawData);
const pantryCopy = clonePantry(pantry);
const actions = planRestock(pantryCopy, shipment);
const grouped = groupByZone(actions);

// Show items expanded in zones
console.log(grouped);

```
