---
id: 68e6a99a0fd7386fb66e156f
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將透過建置一個遊戲設定面板來練習樣式化核取方塊。

所有的 HTML 標準範本（`DOCTYPE`、`html`、`head` 和 `body`）都已為你提供。

在 `body` 內，創建一個包含巢狀 `h1` 元素的 `div` 元素。將文字 `Game Settings` 指定給 `h1` 元素。

# --hints--

你應該有一個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

你應該有一個 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你的 `h1` 元素應該在你的 `div` 元素內。

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

你的 `h1` 元素應該有文字 `Game Settings`。

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
