---
id: 68e9a834dfef752dc35f6f90
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

現在當使用者點擊 `checkboxes` 時會改變顏色，你將會套用一些樣式，使被選取的 `checkbox` 更加明顯。

CSS 偽類別可以串鏈。你可以將一個偽類別或偽元素添加到另一個之後，以標的符合多個條件的元素。

為 `input[type="checkbox"]` 創建一個新的選擇器。為你的選擇器添加 `:checked` 偽類，接著添加 `::after` 偽元素。這讓你能在核取方塊被選取時添加視覺指示器。

接著，為你的 `content` 添加值為 `"✓"`。

# --hints--

你應該有一個結合型別選擇器與偽類別選擇器和偽元素選擇器的 `input[type="checkbox"]:checked::after`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

你應該為你的組合型別選擇器添加一個 `content` 屬性，並搭配偽類選擇器和偽元素選擇器 `input[type="checkbox"]:checked::after`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

你應該有一個 `content` 屬性，值為 `"✓"`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

