---
id: 68e9a54dd97fd9091e18acce
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

現在創建一個 `h1` 選擇器，並將 `text-align` 屬性設為 `center`。這將使你的 `h1` 元素置中。

# --hints--

你應該在你的 CSS 中為 `h1` 元素添加型別選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1"));
```

你應該為你的 `h1` 選擇器添加 `text-align` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign);
```

你應該有一個 `text-align` 屬性，其值為 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

--fcc-editable-region--

--fcc-editable-region--

```
