---
id: 68e9a7bd60c971278d2d4ffc
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

在之前的課程中，你學習了 CSS 中的偽類和偽元素。你現在將運用這些知識，透過建立結合型別選擇器和偽類選擇器來實作。

首先，從型別選擇器的格式開始，將它設定為 `input[type="checkbox"]` 並在其末尾附加 `:checked`。

接著，給它一個值為 `#f1be32` 的 `background-color`。

最後，將 `border-color` 設定為 `#e2a60d`。

# --hints--

你應該有一個結合型別和偽類別選擇器的 `input[type="checkbox"]:checked`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

你應該為你的結合型別和偽類別選擇器 `input[type="checkbox"]:checked` 添加 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

你應該有一個 `background-color` 屬性，值為 `#f1be32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

你應該為你的結合型別和偽類別選擇器 `input[type="checkbox"]:checked` 添加 `border-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

你應該有一個 `border-color` 屬性，值為 `#e2a60d`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
