---
id: 68e9a73075dc26207320811a
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

現在你可以再次看到你的核取方塊，接著你將為它完成一些樣式選項。為你的 `input[type="checkbox"]` 選擇器添加 `4px` 的 `border-radius`，使其邊緣變圓。

然後給它一個 `background-color` 為 `white`，使你的核取方塊中心從容器的背景中突出。

由於你將為使用者點擊核取方塊時設定訂製的過渡效果，請設定一個值為 `all` 和 `0.3s` 的 `transition`，使過渡在 0.3 秒內平滑發生，而非瞬間完成。

# --hints--

你應該為你的 `input[type="checkbox"]` 選擇器添加 `border-radius` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius);
```

你應該有一個 `border-radius` 屬性，值為 `4px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.borderRadius, "4px");
```

你應該為你的 `input[type="checkbox"]` 選擇器添加 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor);
```

你應該有一個 `background-color` 屬性，值為 `white`。

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.backgroundColor, ["white", "rgb(255,255,255)", "#ffffff"]);
```

你應該為你的 `input[type="checkbox"]` 選擇器添加 `transition` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.transition);
```

你應該有一個值為 `all 0.3s` 的 `transition`。

```js
/**
 * Because CSS is simply the most lovely thing to ever exist, the behaviour of the `all` shorthand does wonky things and will not show up
 * in the computed stylesheet. Instead, we need to query the actual style element and extract the selector, checking the properties as written
 * by the camper.
 * I really hate CSS.
 * <3 Naomi
 */
const { props } = document.querySelector('.fcc-injected-styles')?.innerText?.match(/input\[type=('|")checkbox\1\]\s*\{(?<props>.*)\}/s)?.groups;
assert.match(props, /transition\s*\:\s*all\s+0.3\s*/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
