---
id: 68e9a6bf1a74721a4e43e217
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

現在你將移除瀏覽器套用的預設勾選框。在你的 `input[type='checkbox']` 選擇器內，添加值為 `none` 的 `appearance`。

將 `appearance` 屬性設定為 `none` 將會清除瀏覽器套用於核取方塊的外觀，讓你能顯示你想要的樣式。

完成後，由於核取方塊將不再可見，請設定一個 `2px` `solid` 且十六進位色碼為 `#f1be32` 的 `border`。

# --hints--

你應該為你的 `input[type="checkbox"]` 選擇器添加一個 `appearance` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

你應該有一個值為 `none` 的 `appearance` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

你應該為你的 `input[type="checkbox"]` 選擇器添加 `border` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

你應該有一個 `border` 屬性，值為 `2px solid #f1be32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
