---
id: 68e9a6201e2efa142d62bfe1
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

接下來，你將放大核取方塊以提升可見度。

首先設定一個 `input` 的選擇器，但特別針對你的 `[type="checkbox"]`。

在你的選擇器中，將 `width` 和 `height` 設定為 `20px`。這會使它比之前更大。

並且為了符合你在標籤上設定的 `cursor` 設定，為你的 `cursor` 添加並指定值為 `pointer`。之後，當你暫留在核取方塊上時，將會顯示指標。

# --hints--

你的 CSS 中應該有一個 `input[type="checkbox"]` 的選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

你應該為你的 `input[type="checkbox"]` 選擇器添加 `width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

你應該有一個 `width` 屬性，值為 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

你應該為你的 `input[type="checkbox"]` 選擇器添加 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

你應該有一個 `height` 屬性，值為 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

你應該為你的 `input[type="checkbox"]` 選擇器添加一個 `cursor` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

你應該有一個 `cursor` 屬性，其值為 `pointer`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
