---
id: 68e9a3194160f8f5620ddf21
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `h1` 元素之後立即創建四個 `label` 元素。

在每個 `label` 元素內，創建一個 `input` 元素，並將 `type` 屬性設為 `"checkbox"`。

在每個 `input` 元素之後，依序為你的 `label` 元素提供以下文字：`Sound Effects`、`Background Music`、`Hard Mode` 和 `Haptic Feedback`。

# --hints--

你應該有四個 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

你的第一個 `label` 元素應該有一個 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

你的第一個 `input` 元素應該將 `type` 屬性設為 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

你的第一個 `label` 元素應該有文字 `Sound Effects`。

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

你的 `Sound Effects` 文字應該放在 `input` 元素之後。

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

你的第二個 `label` 元素應該有一個 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

你的第二個 `input` 元素的 `type` 屬性應設為 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

你的第二個 `label` 元素應該有文字 `Background Music`。

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

你的 `Background Music` 文字應該放在 `input` 元素之後。

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

你的第三個 `label` 元素應該有一個 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

你的第三個 `input` 元素的 `type` 屬性應設為 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

你的第三個 `label` 元素應該有文字 `Hard Mode`。

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

你的 `Hard Mode` 文字應該放在 `input` 元素之後。

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

你的第四個 `label` 元素應該有一個 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

你的第四個 `input` 元素的 `type` 屬性應設為 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

你的第四個 `label` 元素應該有文字 `Haptic Feedback`。

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

你的 `Haptic Feedback` 文字應該放在 `input` 元素之後。

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
