---
id: 68ee86e3ee07b01626a392cc
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 `.settings-card` 選擇器中將 `margin` 屬性設為 `auto`。

將 `margin` 屬性設定為 `auto` 會自動調整元素的邊距，以平均分配其容器中剩餘的空間，通常用於水平置中區塊層級的元素。

最後，設定一個 `text-align` 屬性，值為 `left`。這將使行內內容，例如文字，對齊到其包含元素的左側。

# --hints--

你應該為你的 `settings-card` 類別選擇器添加 `margin` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

你應該有一個值為 `auto` 的 `margin` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

你應該為你的 `settings-card` 類別選擇器添加一個 `text-align` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

你應該有一個 `text-align` 屬性，值為 `left`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
