---
id: 68f15cd397503429f0345262
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

在 `input[type="checkbox"]:checked::after` 選擇器的宣告中，將 `display` 設定為 `block`。將你的 `display` 屬性設定為 `block` 會使該元素成為區塊層級元素，這表示它會佔據其容器的整個寬度並從新行開始，讓你更容易控制其寬度、高度和間距。

接著，將 `text-align` 設為 `center`。這會將行內內容（例如文字或行內元素）水平置中於區塊內。區塊仍會佔據其容器的整個寬度，但其中的所有內容都會對齊到中間。

# --hints--

你應該為你的組合型別選擇器添加一個 `display` 屬性，該選擇器帶有偽類選擇器和偽元素選擇器 `input[type="checkbox"]:checked::after`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

你應該有一個 `display` 屬性，其值為 `block`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

你應該為你的結合型別選擇器添加一個帶有偽類選擇器和偽元素選擇器 `input[type="checkbox"]:checked::after` 的 `text-align` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

你應該有一個 `text-align` 屬性，其值為 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
