---
id: 68f1560aff1214103ab6b59b
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在你將開始用一些 CSS 樣式美化頁面。先從建立 `body` 選擇器開始。

將 `body` 設定為具有值為 `100vh` 的 `height` 屬性和值為 `#f0f0f0` 的 `background-color` 屬性。

`height` 設為 `100vh` 使 `body` 佔滿瀏覽器視窗的全部高度，而淺灰色背景色則為頁面提供細膩的基底。

最後，設定一個 `text-align` 屬性，值為 `center`。這將置中頁面中所有行內內容，除非子元素以自己的對齊方式改寫它。

# --hints--

你應該使用 `body` 型別選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

你應該在 `body` 類型選擇器中添加 `height` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

你應該有一個 `height` 屬性，值為 `100vh`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

你應該在 `body` 型別選擇器中添加 `background-color` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

你應該有一個 `background-color` 屬性，其值為 `#f0f0f0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

你應該在 `body` 型別選擇器中添加 `text-align` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

你應該有一個 `text-align` 屬性，其值為 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
