---
id: 68f15ef0446464325ce1aa45
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

你需要在核取方塊和標籤之間留一些間距。首先在你的 CSS 中使用 `label` 選擇器，並設定 `display` 屬性為 `block`。

接著，設定 `margin` 屬性為 `8px auto`，以在元素之間添加垂直間距。這將在每個元素的上下方產生一致的間距，有助於分隔內容並提升可讀性。

最後，將 `cursor` 屬性設定為 `pointer` 在 `label` 元素上。當使用者暫留在 `label` 上時，這將把游標變成手形圖示，表示該元素可點擊並提升整體使用者體驗。

# --hints--

你的 CSS 中應該有一個 `label` 的型別選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

你應該為你的 `label` 選擇器添加 `display` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

你應該有一個 `display` 屬性，其值為 `block`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

你應該為你的 `label` 選擇器添加 `margin` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

你應該有一個 `margin` 屬性，值為 `8px auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

你應該為你的 `label` 選擇器添加 `cursor` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

你應該有一個 `cursor` 屬性，其值為 `pointer`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
