---
id: 68f29dbf9125476b4808d6fb
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

現在是時候為 `settings-card` 容器設計樣式。為 `settings-card` 創建一個類別選擇器。

這裡將放置你所有針對容器的格式設定。將 `max-width` 設為 `250px` 以定義你的容器整體大小。

接著，將你的 `padding` 設為 `20px`，以便你的內容與容器的邊框之間有空間。

之後，透過將你的 `border-radius` 設定為 `10px` 來創建圓角邊框。

然後設定一個 `box-shadow`，其值為 `0 2px 6px rgba(0,0,0,0.2)`。這將產生一個細微的「浮起」效果，為容器創造深度。

# --hints--

你的 css 中應該有一個用於 `settings-card` 的類別選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

你應該為你的 `.settings-card` 類別選擇器添加 `max-width` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

你應該有一個 `max-width` 屬性，值為 `250px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

你應該為 `.settings-card` 類別選擇器添加 `padding` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

你應該有一個 `padding` 屬性，值為 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

你應該為你的 `.settings-card` 類別選擇器添加 `border-radius` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

你應該有一個 `border-radius` 屬性，值為 `10px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

你應該為你的 `.settings-card` 類別選擇器添加 `box-shadow` 屬性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

你應該有一個值為 `0 2px 6px rgba(0,0,0,0.2)` 的 `box-shadow`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
