---
id: a3bfc1673c0526e06d3ac698
title: 建置一個質數總和計算器
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

在這個實驗中，你將建置一個計算機，該計算機接受一個數字並傳回所有小於或相等於該數字的質數總和。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個 `sumPrimes` 函式（程式），它接受一個數字作為引數。
1. `sumPrimes` 函式（程式）應該傳回小於或等於所提供數字的所有質數的總和。
1. 如果輸入數字小於 `2`，函式應該傳回 `0`。

# --hints--

你應該有一個 `sumPrimes` 函式（程式）。

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` 應該傳回 `17`。

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` 應該傳回 `10`。

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` 應該傳回 `2`。

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` 應該傳回 `0`。

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` 應該傳回 `73156`。

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
