---
id: 66d75dd0aa65a71600dc669b
title: "Costruisci un programma per la gestione dell'inventario"
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

In questo laboratorio costruirai un programma per la gestione dell'inventario che ti permetterà di aggiungere, aggiornare, trovare e rimuovere prodotti dall'inventario. Userai un array di oggetti per rappresentare il tuo inventario, dove ogni oggetto avrà come chiavi nome e quantità.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti dichiarare un array vuoto chiamato `inventory` che archivierà oggetti prodotto aventi una chiave `name` con il valore di una stringa in minuscolo, e una chiave `quantity` con il valore di un intero.
1. Dovresti creare una funzione chiamata `findProductIndex` che prende come argomento il nome del prodotto e restituisce l'indice dell'oggetto prodotto corrispondente all'interno dell'array `inventory`. La funzione dovrebbe sempre usare la forma minuscola del nome del prodotto per effettuare la ricerca. Se il prodotto non viene trovato, la funzione dovrebbe restituire `-1`.
1. Dovresti creare una funzione chiamata `addProduct` che prende come argomento un oggetto prodotto.
1. Se il prodotto è già presente nell'inventario, la funzione `addProduct` dovrebbe aggiornare il valore `quantity` aggiungendo la quantità passata alla funzione alla quantità `quantity` attuale e stampare sulla console il nome del prodotto seguito da uno spazio e `quantity updated`.
1. Se il prodotto non è presente nell'inventario, la funzione `addProduct` dovrebbe aggiungere il prodotto all'array `inventory` e stampare sulla console il nome del prodotto seguito da uno spazio e `added to inventory`.
1. Dovresti creare una funzione chiamata `removeProduct` che prende come argomenti il nome del prodotto e la quantità.
1. La funzione `removeProduct` dovrebbe sottrarre la quantità passata dall'oggetto prodotto corrispondente nell'inventario e stampare sulla console la stringa `Remaining <product-name> pieces: <product-quantity>`, dove `<product-name>` dovrebbe essere sostituito dal nome del prodotto e `<product-quantity>` dalla quantità del prodotto.
1. Se la quantità dopo la sottrazione è zero, `removeProduct` dovrebbe rimuovere l'oggetto prodotto dall'inventario. Se la quantità nell'inventario non è sufficiente per effettuare la sottrazione, la funzione `removeProduct` dovrebbe stampare sulla console la stringa `Not enough <product-name> available, remaining pieces: <product-quantity>`.
1. Se il prodotto da rimuovere non è presente nell'inventario, la funzione `removeProduct` dovrebbe stampare `<product-name> not found` sulla console.

**Nota:** Per evitare conflitti, mantieni solo le stampe sulla console menzionate nelle storie utente quando esegui i test.

# --hints--

Dovresti dichiarare un array vuoto chiamato `inventory`.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

Dovresti avere una funzione chiamata `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` dovrebbe restituire l'indice dell'oggetto che ha `name` uguale a `"flour"` all'interno dell'array `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` dovrebbe restituire l'indice dell'oggetto che ha `name` uguale a `"flour"` all'interno dell'array `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` dovrebbe restituire `-1` quando nessun oggetto con `name` uguale a `"flour"` viene trovato all'interno dell'array `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

Dovresti avere una funzione chiamata `addProduct`.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` dovrebbe aggiungere `5` al valore `quantity` dell'oggetto che ha `name` uguale a `"flour"` nell'array `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` dovrebbe stampare `flour quantity updated` quando un oggetto con `name` uguale a `"flour"` viene trovato nell'array `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` dovrebbe aggiungere `{name: "flour", quantity: 5}` all'array `inventory` quando nessun oggetto con `name` uguale a `"flour"` viene trovato nell'inventario.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` dovrebbe stampare `flour added to inventory` quando nessun oggetto con `name` uguale a `"flour"` viene trovato nell'inventario.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

Dovresti creare una funzione chiamata `removeProduct`.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` dovrebbe stampare `flour not found` quando nessun oggetto con `name` uguale a `"flour"` viene trovato nell'array `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` dovrebbe sottrarre `5` dal valore `quantity` dell'oggetto che ha `name` uguale a `"flour"` all'interno dell'array `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` dovrebbe stampare `Remaining flour pieces: 15` sulla console, quando `inventory` è uguale a `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

Se la `quantity` dopo la sottrazione è zero, `removeProduct` dovrebbe rimuovere l'oggetto prodotto dall'inventario.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` dovrebbe stampare `Not enough flour available, remaining pieces: 5` quando `inventory` è uguale a `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
