---
id: 587d7b85367417b2b2512b3a
title: 関数呼び出し時に引数が間違った順序で渡された場合の検出
challengeType: 1
forumTopicId: 301184
dashedName: catch-arguments-passed-in-the-wrong-order-when-calling-a-function
---

# --description--

関数の呼び出しに関する話を続けます。次に注意すべきバグは、関数の引数が間違った順序で渡される場合です。引数の型が異なる場合、例えば関数が配列と整数を期待しているのに順序が逆だと、実行時エラーが発生する可能性があります。引数の型が同じ場合（例えばすべて整数の場合）は、コードのロジックが意味をなさなくなります。これらの問題を避けるために、すべての必要な引数を正しい順序で渡すようにしてください。

# --instructions--

関数`raiseToPower`は底を指数に上げます。残念ながら、正しく呼び出されていません。`power`の値が期待される8になるようにコードを修正してください。

# --hints--

`power`の変数を、3の2乗ではなく2の3乗になるように修正してください。

```js
assert(power == 8);
```

`raiseToPower`関数呼び出しの引数の順序を正しく使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/raiseToPower\(\s*?base\s*?,\s*?exp\s*?\);/g));
```

# --seed--

## --seed-contents--

```js
function raiseToPower(b, e) {
  return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(exp, base);
console.log(power);
```

# --solutions--

```js
function raiseToPower(b, e) {
 return Math.pow(b, e);
}

let base = 2;
let exp = 3;
let power = raiseToPower(base, exp);
console.log(power);
```
