---
id: 587d7b86367417b2b2512b3d
title: 有効な終了条件で無限ループを防ぐ
challengeType: 1
forumTopicId: 301192
dashedName: prevent-infinite-loops-with-a-valid-terminal-condition
---

# --description--

最後のトピックは恐ろしい無限ループです。ループは、プログラムでコードブロックを特定の回数だけ、または条件が満たされるまで実行したいときに便利なツールですが、ループを終了させる終了条件が必要です。無限ループはブラウザをフリーズさせたりクラッシュさせたり、プログラムの実行を混乱させる可能性があり、誰も望みません。

このセクションの導入部分に無限ループの例がありましたが、`while` ループ内の `loopy()` に終了条件がありませんでした。この関数を呼び出さないでください！

```js
function loopy() {
  while(true) {
    console.log("Hello, world!");
  }
}
```

終了条件は、プログラムにループコードから抜けるタイミングを伝えるもので、最終的に到達する必要があるのはプログラマーの役目です。よくあるミスは、終了条件から見てカウンター変数を間違った方向に増減させることです。もう一つは、ループコード内でカウンターやインデックス変数を増減させる代わりに誤ってリセットしてしまうことです。

# --instructions--

`myFunc()` 関数には無限ループがあります。終了条件 `i != 4` は決して `false` と評価されず（ループを抜けません）、`i` は毎回2ずつ増加し、`i` が奇数で始まるため4を飛び越えてしまいます。終了条件の比較演算子を修正して、ループが `i` が4以下の間だけ実行されるようにしてください。

# --hints--

終了条件（中央部分）の比較演算子を `for` ループ内で変更してください。

```js
assert(__helpers.removeJSComments(code).match(/i\s*?<=\s*?4;/g).length == 1);
```

ループの終了条件の比較演算子を修正してください。

```js
assert(!__helpers.removeJSComments(code).match(/i\s*?!=\s*?4;/g));
```

# --seed--

## --seed-contents--

```js
function myFunc() {
  for (let i = 1; i != 4; i += 2) {
    console.log("Still going!");
  }
}
```

# --solutions--

```js
function myFunc() {
 for (let i = 1; i <= 4; i += 2) {
   console.log("Still going!");
 }
}
```
