---
id: 587d7b84367417b2b2512b36
title: '開いている丸括弧、括弧、波括弧、引用符の閉じ忘れを見つける'
challengeType: 1
forumTopicId: 301190
dashedName: catch-unclosed-parentheses-brackets-braces-and-quotes
---

# --description--

注意すべきもう一つの構文エラーは、すべての開き丸括弧、括弧、波括弧、引用符に対応する閉じるペアがあることです。ペアの一方を忘れてしまうのは、既存のコードを編集しているときに、そのペアのどちらかを挿入する場合に起こりやすいです。また、メソッドの引数としてコールバック関数を追加するなど、コードブロックを他の中にネストするときも注意してください。

このミスを避ける方法の一つは、開き文字を入力したらすぐに対応する閉じ文字も入力し、カーソルをその間に戻してからコーディングを続けることです。幸いなことに、ほとんどの最新のコードエディターはペアの後半を自動生成してくれます。

# --instructions--

コード内の2つのペアのエラーを修正してください。

# --hints--

配列の欠けている部分を修正してください。

```js
assert(__helpers.removeJSComments(code).match(/myArray\s*?=\s*?\[\s*?1\s*?,\s*?2\s*?,\s*?3\s*?\];/g));
```

`.reduce()` メソッドの欠けている部分を修正してください。コンソールの出力は `Sum of array values is: 6` であるべきです。

```js
assert(arraySum === 6);
```

# --seed--

## --seed-contents--

```js
let myArray = [1, 2, 3;
let arraySum = myArray.reduce((previous, current =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```

# --solutions--

```js
let myArray = [1, 2, 3];
let arraySum = myArray.reduce((previous, current) =>  previous + current);
console.log(`Sum of array values is: ${arraySum}`);
```
