---
id: 587d7b85367417b2b2512b38
title: 代入演算子の代わりに等値演算子の使用を検出する
challengeType: 1
forumTopicId: 301191
dashedName: catch-use-of-assignment-operator-instead-of-equality-operator
---

# --description--

分岐処理、つまり特定の条件が満たされた場合に異なる処理を行うプログラムは、JavaScriptの`if`、`else if`、および`else`文に依存しています。条件は結果がある値と等しいかどうかをテストする形をとることがあります。

このロジックは（少なくとも英語では）「if x equals y, then ...」と表現され、これは文字通り`=`、つまり代入演算子を使ったコードに翻訳されることがあります。これによりプログラムの制御フローが予期せぬものになります。

前のチャレンジで説明したように、JavaScriptの代入演算子（`=`）は変数名に値を代入します。そして`==`および`===`演算子は等値をチェックします（三重の`===`は厳密な等値をテストし、値と型の両方が同じであることを意味します）。

以下のコードは`x`に2を代入しており、これは`true`として評価されます。JavaScriptではほとんどすべての値は単独で`true`として評価されますが、「falsy」と呼ばれる値は例外で、`false`、`0`、`""`（空文字列）、`NaN`、`undefined`、および`null`です。

```js
let x = 1;
let y = 2;
if (x = y) {

} else {

}
```

この例では、`if`がfalsyでない限り、`y`のどんな値でも`y`文内のコードブロックが実行されます。ここで実行されることを期待している`else`ブロックは実際には実行されません。

# --instructions--

条件を修正してプログラムが正しい分岐を実行し、適切な値が`result`に代入されるようにしてください。

# --hints--

条件を代入ではなく等値のチェックに修正してください。

```js
assert(result == 'Not equal!');
```

条件は等値をテストするために`==`または`===`のいずれかを使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/x\s*?===?\s*?y/g));
```

# --seed--

## --seed-contents--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x = y) {
  result = "Equal!";
} else {
  result = "Not equal!";
}

console.log(result);
```

# --solutions--

```js
let x = 7;
let y = 9;
let result = "to come";

if(x === y) {
 result = "Equal!";
} else {
 result = "Not equal!";
}

console.log(result);
```
