---
id: 587d7b85367417b2b2512b39
title: 関数呼び出し後の開き丸括弧と閉じ丸括弧の抜けを検出する
challengeType: 1
forumTopicId: 301185
dashedName: catch-missing-open-and-closing-parenthesis-after-a-function-call
---

# --description--

関数やメソッドが引数を取らない場合、呼び出すときに（空の）開き丸括弧と閉じ丸括弧を入れ忘れることがあります。多くの場合、関数呼び出しの結果は変数に保存され、コード内の他の場所で使われます。このエラーは、変数の値（または型）をコンソールにログ出力して、期待される関数の戻り値ではなく関数の参照が設定されていることを確認することで検出できます。

以下の例では変数の中身が異なります：

```js
function myFunction() {
  return "You rock!";
}
let varOne = myFunction;
let varTwo = myFunction();
```

ここで `varOne` は関数 `myFunction` であり、`varTwo` は文字列 `You rock!` です。

# --instructions--

変数 `result` に関数 `getNine` を呼び出した結果の値が設定されるようにコードを修正してください。

# --hints--

変数 `result` が関数 `getNine` が返す数値に設定されるようにコードを修正してください。

```js
assert(result == 9);
```

`getNine` 関数を呼び出すコードにしてください。

```js
assert(__helpers.removeJSComments(code).match(/getNine\(\)/g).length == 2);
```

# --seed--

## --seed-contents--

```js
function getNine() {
  let x = 6;
  let y = 3;
  return x + y;
}

let result = getNine;
console.log(result);
```

# --solutions--

```js
function getNine() {
 let x = 6;
 let y = 3;
 return x + y;
}

let result = getNine();
console.log(result);
```
