---
id: 587d7b83367417b2b2512b33
title: JavaScriptコンソールを使って変数の値を確認する
challengeType: 1
forumTopicId: 18372
dashedName: use-the-javascript-console-to-check-the-value-of-a-variable
---

# --description--

ChromeとFirefoxには、JavaScriptのデバッグに役立つ優れたJavaScriptコンソール（開発者ツールとも呼ばれます）が備わっています。

ChromeのメニューやFirefoxのメニューから開発者ツールやウェブコンソールを見つけることができます。別のブラウザや携帯電話を使っている場合は、デスクトップ版のFirefoxかChromeに切り替えることを強くお勧めします。

`console.log()` メソッドは、丸括弧内の内容をコンソールに「出力する」ため、最も役立つデバッグツールになるでしょう。コードの要所に配置することで、変数の途中の値を確認できます。出力がどうなるべきかをあらかじめイメージしておくのが良い習慣です。コード全体の計算状況をチェックポイントで確認することで、問題の箇所を絞り込むのに役立ちます。

以下は文字列 `Hello world!` をコンソールに出力する例です。

```js
console.log('Hello world!');
```

# --instructions--

コード内の指定された場所で、変数 `console.log()` の値を`a` メソッドを使って出力してください。

# --hints--

コードでは、変数 `console.log()` の値を確認するために`a` を使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/console\.log\(a\)/g));
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 1;
a++;
// Only change code below this line


let sumAB = a + b;
console.log(sumAB);
```

# --solutions--

```js
var a = 5; console.log(a);
```
