---
id: 587d7b86367417b2b2512b3c
title: ループ内で変数を再初期化するときは注意してください
challengeType: 1
forumTopicId: 301194
dashedName: use-caution-when-reinitializing-variables-inside-a-loop
---

# --description--

ループ内で情報を保存したり、カウンターを増やしたり、変数をリセットしたりする必要がある場合があります。問題になるのは、変数を再初期化すべきなのにしなかったり、その逆であったりする場合です。特に、終了条件に使われている変数を誤ってリセットしてしまうと、無限ループになる危険があります。

ループの各サイクルで変数の値を`console.log()`を使って出力すると、変数のリセットやリセット忘れに関するバグの原因を見つけやすくなります。

# --instructions--

次の関数は、`m`行`n`列のゼロの二次元配列を作成するはずです。残念ながら、外側のループで`row`変数が再初期化（空の配列に戻す）されていないため、期待した出力が得られていません。正しい3×2のゼロの配列（`[[0, 0], [0, 0], [0, 0]]`のようなもの）を返すようにコードを修正してください。

# --hints--

コードは`matrix`変数を、3行2列のゼロを持つ配列に設定する必要があります。

```js
assert(JSON.stringify(matrix) == '[[0,0],[0,0],[0,0]]');
```

`matrix`変数は3行であるべきです。

```js
assert(matrix.length == 3);
```

`matrix`変数は各行に2列あるべきです。

```js
assert(
  matrix[0].length == 2 && matrix[1].length === 2 && matrix[2].length === 2
);
```

`zeroArray(4,3)`は4行3列のゼロを持つ配列を返すべきです。

```js
assert(JSON.stringify(zeroArray(4,3)) == '[[0,0,0],[0,0,0],[0,0,0],[0,0,0]]');
```

# --seed--

## --seed-contents--

```js
function zeroArray(m, n) {
  // Creates a 2-D array with m rows and n columns of zeroes
  let newArray = [];
  let row = [];
  for (let i = 0; i < m; i++) {
    // Adds the m-th row into newArray

    for (let j = 0; j < n; j++) {
      // Pushes n zeroes into the current row to create the columns
      row.push(0);
    }
    // Pushes the current row, which now has n zeroes in it, to the array
    newArray.push(row);
  }
  return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```

# --solutions--

```js
function zeroArray(m, n) {
 // Creates a 2-D array with m rows and n columns of zeroes
 let newArray = [];
 for (let i = 0; i < m; i++) {
   let row = [];
   // Adds the m-th row into newArray

   for (let j = 0; j < n; j++) {
     // Pushes n zeroes into the current row to create the columns
     row.push(0);
   }
   // Pushes the current row, which now has n zeroes in it, to the array
   newArray.push(row);
 }
 return newArray;
}

let matrix = zeroArray(3, 2);
console.log(matrix);

```
