---
id: ab306dbdcc907c7ddfc30830
title: створити інструмент для глибокого згладжування
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм згладжування масиву.

Згладжування масиву означає перетворення вкладеного масиву будь-якої глибини на один одномірний масив. Процес витягує всі елементи послідовно, розгортаючи лише масиви. Інші типи залишаються без змін.

Наприклад:

| Оригінал | Згладжений |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `steamrollArray`.
1. Функція `steamrollArray` повинна приймати один аргумент: вкладений масив.
1. Функція має згладжувати вкладений масив, враховуючи різні рівні вкладеності.
1. Ваше рішення не повинно використовувати методи `Array.prototype.flat()` або `Array.prototype.flatMap()`.
1. Не слід використовувати глобальні змінні.

# --hints--

У вас має бути функція `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` має повертати `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` має повертати `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` має повертати `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` має повертати `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Ваше рішення не повинно використовувати методи `Array.prototype.flat()` або `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Не використовуйте глобальні змінні.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
