---
id: 66d75dd0aa65a71600dc669b
title: 在庫管理プログラムを作成する
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

このラボでは、在庫から商品を追加、更新、検索、削除できる在庫管理プログラムを作成します。各オブジェクトが name と quantity をキーに持つオブジェクトの配列を使って在庫を表現します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストを通過してラボを完了してください。

**ユーザーストーリー:**

1. `inventory` という空の配列を宣言し、商品オブジェクトを格納してください。商品オブジェクトは `name` キーに小文字の文字列の値を持ち、`quantity` キーに整数の値を持ちます。
1. `findProductIndex` という関数を作成してください。この関数は商品名を引数に取り、`inventory` 配列内の対応する商品オブジェクトのインデックスを返します。検索には常に商品名の小文字を使ってください。商品が見つからない場合は `-1` を返します。
1. `addProduct` という関数を作成してください。この関数は商品オブジェクトを引数に取ります。
1. 商品がすでに在庫にある場合、`addProduct` 関数はその商品の `quantity` の値を、関数に渡された数量を現在の `quantity` に加算して更新し、商品名とスペース、`quantity updated` をコンソールに出力してください。
1. 商品が在庫にない場合、`addProduct` 関数は商品を `inventory` 配列に追加し、商品名とスペース、`added to inventory` をコンソールに出力してください。
1. `removeProduct` という関数を作成してください。この関数は商品名と数量を引数に取ります。
1. `removeProduct` 関数は渡された数量を在庫内の対応する商品オブジェクトから減算し、文字列 `Remaining <product-name> pieces: <product-quantity>` をコンソールに出力してください。この文字列では `<product-name>` が商品名に、`<product-quantity>` が商品数量に置き換えられます。
1. 減算後の数量がゼロの場合、`removeProduct` は商品オブジェクトを在庫から削除します。在庫の数量が減算に足りない場合、`removeProduct` 関数は文字列 `Not enough <product-name> available, remaining pieces: <product-quantity>` をコンソールに出力してください。
1. 削除対象の商品が在庫にない場合、`removeProduct` 関数は `<product-name> not found` をコンソールに出力してください。

**注意:** テスト実行時にはユーザーストーリーに記載されたログ出力のみを行い、他のログは出さないようにしてください。

# --hints--

`inventory` という空の配列を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

`findProductIndex` という関数を用意してください。

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` は `name` 配列内で `"flour"` が `inventory` と等しいオブジェクトのインデックスを返すべきです。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` は `name` 配列内で `"flour"` が `inventory` と等しいオブジェクトのインデックスを返すべきです。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` は `-1` 配列内で `name` が `"flour"` と等しいオブジェクトが見つからない場合、`inventory` を返すべきです。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

`addProduct` という関数を用意してください。

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` は `5` 配列内で `quantity` が `name` と等しいオブジェクトの `"flour"` の値に `inventory` を加算すべきです。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` は `flour quantity updated` 配列内で `name` が `"flour"` と等しいオブジェクトが見つかった場合、`inventory` をログに出力すべきです。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` は 在庫に `{name: "flour", quantity: 5}` が `inventory` と等しいオブジェクトが見つからない場合、`name` を `"flour"` 配列に追加すべきです。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` は 在庫に `flour added to inventory` が `name` と等しいオブジェクトが見つからない場合、`"flour"` をログに出力すべきです。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

`removeProduct` という関数を作成してください。

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` は `flour not found` 配列内で `name` が `"flour"` と等しいオブジェクトが見つからない場合、`inventory` をログに出力すべきです。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` は `5` 配列内で `quantity` が `name` と等しいオブジェクトの `"flour"` の値から `inventory` を減算すべきです。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` は `Remaining flour pieces: 15` が `inventory` と等しい場合、コンソールに `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]` を出力すべきです。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

減算後の `quantity` がゼロの場合、`removeProduct` は商品オブジェクトを在庫から削除すべきです。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` は `Not enough flour available, remaining pieces: 5` が `inventory` と等しい場合、`[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]` をログに出力すべきです。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
