---
id: 6140cdebd39d6a101e747529
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

`.line` 要素用のセレクターを作成してください。 まずは `background-color` を `black` に、`width` を `50%` に、`height` を `2px` に設定してください。

# --hints--

You should have a `.line` selector.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

Your `.line` selector should have a `background-color` property set to `black`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

Your `.line` selector should have a `width` property set to `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

Your `.line` selector should have a `height` property set to `2px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
