---
id: 587d7fb2367417b2b2512bf5
title: "Ottenere l'input del parametro di rotta dal client"
challengeType: 2
forumTopicId: 301513
dashedName: get-route-parameter-input-from-the-client
---

# --description--

Quando costruisci un'API, devi permettere agli utenti di comunicarti cosa vogliono ottenere dal tuo servizio. Per esempio, se il client richiede informazioni su un utente memorizzato nel database, ha bisogno di un modo per farti sapere quale utente gli interessa. Un modo possibile per ottenere questo risultato è usare i parametri di rotta. I parametri di rotta sono segmenti nominati dell'URL, delimitati da slash (/). Ogni segmento cattura il valore della parte dell'URL che corrisponde alla sua posizione. I valori catturati si trovano nell'oggetto `req.params`.

<blockquote>route_path: '/user/:userId/book/:bookId'<br>actual_request_URL: '/user/546/book/6754' <br>req.params: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Costruisci un server echo, montato sulla rotta `GET /:word/echo`. Rispondi con un oggetto JSON, seguendo la struttura `{echo: word}`. Puoi trovare la parola da ripetere in `req.params.word`. Puoi testare la tua rotta dalla barra degli indirizzi del browser, visitando alcune rotte corrispondenti, per esempio `your-app-rootpath/freecodecamp/echo`.

# --hints--

Test 1 : Il tuo server echo dovrebbe ripetere correttamente le parole

```js
  const response = await fetch(code + '/eChOtEsT/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'eChOtEsT',
    'Test 1: the echo server is not working as expected'
  );
```

Test 2 : Il tuo server echo dovrebbe ripetere correttamente le parole

```js
  const response = await fetch(code + '/ech0-t3st/echo');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.echo,
    'ech0-t3st',
    'Test 2: the echo server is not working as expected'
  );
```
