---
id: 587d7fb2367417b2b2512bf6
title: Ottenere input da parametri query dal client
challengeType: 2
forumTopicId: 301512
dashedName: get-query-parameter-input-from-the-client
---

# --description--

Un altro modo comune per ottenere input dal client è codificare i dati dopo il percorso della rotta, usando una stringa di query. La stringa di query è delimitata da un punto interrogativo (?), e include coppie campo=valore. Ogni coppia è separata da un e commerciale (&). Express può analizzare i dati dalla stringa di query e popolare l'oggetto `req.query`. Alcuni caratteri, come il percentuale (%), non possono essere presenti negli URL e devono essere codificati in un formato diverso prima di poterli inviare. Se usi l'API di JavaScript, puoi usare metodi specifici per codificare/decodificare questi caratteri.

<blockquote>route_path: '/library'<br>actual_request_URL: '/library?userId=546&#x26;bookId=6754' <br>req.query: {userId: '546', bookId: '6754'}</blockquote>

# --instructions--

Crea un endpoint API, montato su `GET /name`. Rispondi con un documento JSON, seguendo la struttura `{ name: 'firstname lastname'}`. I parametri first e last name devono essere codificati in una stringa di query, ad esempio `?first=firstname&last=lastname`.

**Nota:** Nel seguente esercizio riceverai dati da una richiesta POST, allo stesso percorso di rotta `/name`. Se vuoi, puoi usare il metodo `app.route(path).get(handler).post(handler)`. Questa sintassi ti permette di concatenare diversi gestori di verbo sullo stesso percorso di rotta. Puoi risparmiare un po' di digitazione e avere un codice più pulito.

# --hints--

Test 1 : Il tuo endpoint API dovrebbe rispondere con `{ "name": "Mick Jagger" }` quando l'endpoint `/name` viene chiamato con `?first=Mick&last=Jagger`

```js
  const response = await fetch(code + '/name?first=Mick&last=Jagger');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "GET /name" route does not behave as expected'
  );
```

Test 2 : Il tuo endpoint API dovrebbe rispondere con `{ "name": "Keith Richards" }` quando l'endpoint `/name` viene chiamato con `?first=Keith&last=Richards`

```js
  const response = await fetch(code + '/name?last=Richards&first=Keith');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "GET /name" route does not behave as expected'
  );
```
