---
id: 587d7fb2367417b2b2512bf8
title: Ottenere dati dalle richieste POST
challengeType: 2
forumTopicId: 301511
dashedName: get-data-from-post-requests
---

# --description--

Monta un gestore POST sul percorso `/name`. È lo stesso percorso di prima. Abbiamo preparato un modulo nella pagina principale html. Invierà gli stessi dati dell’esercizio 10 (Query string). Se body-parser è configurato correttamente, dovresti trovare i parametri nell’oggetto `req.body`. Dai un’occhiata all’esempio abituale della libreria:

<blockquote>rotta: POST '/library'<br>urlencoded_body: userId=546&#x26;bookId=6754 <br>req.body: {userId: '546', bookId: '6754'}</blockquote>

Rispondi con lo stesso oggetto JSON di prima: `{name: 'firstname lastname'}`. Verifica se il tuo endpoint funziona usando il modulo html che abbiamo fornito nella pagina principale dell’app.

Suggerimento: ci sono diversi altri metodi http oltre a GET e POST. E per convenzione c’è una corrispondenza tra il verbo http e l’operazione che eseguirai sul server. La mappatura convenzionale è:

POST (a volte PUT) - Creare una nuova risorsa usando le informazioni inviate con la richiesta,

GET - Leggere una risorsa esistente senza modificarla,

PUT o PATCH (a volte POST) - Aggiornare una risorsa usando i dati inviati,

DELETE - Eliminare una risorsa.

Ci sono anche un paio di altri metodi usati per negoziare una connessione con il server. Tranne GET, tutti gli altri metodi elencati sopra possono avere un payload (cioè i dati nel corpo della richiesta). Il middleware body-parser funziona anche con questi metodi.

# --hints--

Test 1 : Il tuo endpoint API dovrebbe rispondere con il nome corretto

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({ first: 'Mick', last: 'Jagger' })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Mick Jagger',
    'Test 1: "POST /name" route does not behave as expected'
  );
```

Test 2 : Il tuo endpoint API dovrebbe rispondere con il nome corretto

```js
  const response = await fetch(code + '/name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/x-www-form-urlencoded'
    },
    body: new URLSearchParams({
      first: 'Keith',
      last: 'Richards'
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.name,
    'Keith Richards',
    'Test 2: "POST /name" route does not behave as expected'
  );
```
