---
id: 587d7fb0367417b2b2512bef
title: Servire un file HTML
challengeType: 2
forumTopicId: 301516
dashedName: serve-an-html-file
---

# --description--

Puoi rispondere alle richieste con un file usando il metodo `res.sendFile(path)`. Puoi inserirlo all’interno del gestore della rotta `app.get('/', ...)`. Dietro le quinte, questo metodo imposta le intestazioni appropriate per indicare al browser come gestire il file che vuoi inviare, in base al suo tipo. Poi legge e invia il file. Questo metodo richiede un percorso assoluto del file. Ti consigliamo di usare la variabile globale di Node `__dirname` per calcolare il percorso in questo modo:

```js
absolutePath = __dirname + '/relativePath/file.ext'
```

# --instructions--

Invia il file `/views/index.html` come risposta alle richieste GET al percorso `/`. Se visualizzi la tua app live, dovresti vedere una grande intestazione HTML (e un modulo che useremo più avanti…), senza alcuno stile applicato.

**Nota:** Puoi modificare la soluzione della sfida precedente o crearne una nuova. Se crei una nuova soluzione, tieni presente che Express valuta le rotte dall’alto verso il basso ed esegue il gestore della prima corrispondenza. Devi commentare la soluzione precedente, altrimenti il server continuerà a rispondere con una stringa.

# --hints--

La tua app dovrebbe servire il file views/index.html

```js
  const response = await fetch(code);
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  assert.match(
    data,
    /<h1>.*<\/h1>/,
    'Your app does not serve the expected HTML'
  );
```
