---
id: 60a3e3396c7b40068ad69970
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

Before you can start styling the `div` you added, you need to link your CSS to your HTML.

添加 `link` 元素以鏈接你的 `styles.css` 文件。 將 `href` 設置爲 `styles.css`，並記住將 `rel` 屬性設置爲 `stylesheet`。

# --hints--

你的代碼應該有一個 `link` 元素。

```js
assert(/<link/.test(code))
```

`link` 元素應該有一個 `rel` 屬性，其值爲 `stylesheet`。

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

`link` 元素應該有一個 `href` 屬性，其值爲 `styles.css`。

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
