---
id: 6976db7ecfa770bf21307b20
title: construir um motor de remix de playlist
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

Neste laboratório, você vai construir um programa que cria uma única playlist de remix a partir de múltiplas playlists enviadas por ouvintes.

Cada ouvinte fornece uma lista de músicas que quer ouvir. Algumas músicas podem aparecer mais de uma vez, e alguns artistas podem aparecer muitas vezes. Seu trabalho é processar essas playlists passo a passo: combiná-las em uma lista só, pontuar cada música, remover músicas duplicadas, limitar quantas vezes o mesmo artista aparece e então criar uma ordem final de reprodução.

**Objetivo:** Atenda às user stories abaixo e faça todos os testes passarem para completar o laboratório.

**User stories:**

1. Você deve criar uma função chamada `flattenPlaylists` que aceita um array de playlists onde cada playlist é um array de objetos com as seguintes propriedades: `trackId`, `artist`, `title`, `votes`, `bpm`. Se a entrada não for um array, `flattenPlaylists` deve retornar um array vazio. Um exemplo de playlist foi fornecido para você. Você pode usar esse exemplo para testar sua função.

2. `flattenPlaylists` deve retornar um array plano de objetos de faixa, onde cada objeto inclui todas as propriedades originais da faixa mais uma propriedade `source` definida como um array com o índice da playlist e o índice da faixa indicando de onde a faixa veio.

3. Você deve criar uma função chamada `scoreTracks` que aceita um array de objetos de faixa conforme retornado por `flattenPlaylists` (cada um com propriedades `trackId`, `artist`, `title`, `votes`, `bpm` e `source`) e retorna um novo array de objetos de faixa, cada um com uma propriedade `score` adicionada usando a fórmula: `votes * 10 - Math.abs(bpm - 120)`.

4. Você deve criar uma função chamada `dedupeTracks` que aceita um array de objetos de faixa conforme retornado por `scoreTracks` e retorna um novo array com entradas duplicadas de `trackId` removidas, mantendo apenas a primeira ocorrência de cada uma.

5. Você deve criar uma função chamada `enforceArtistQuota` que aceita um array de objetos de faixa conforme retornado por `dedupeTracks` e um número representando o máximo de ocorrências permitidas por artista. A função deve retornar um novo array onde nenhum artista aparece mais vezes do que o número dado, mantendo as ocorrências mais antigas.

6. Você deve criar uma função chamada `buildSchedule` que aceita um array de objetos de faixa conforme retornado por `enforceArtistQuota` e retorna um novo array de objetos `{ slot, trackId }`, onde `slot` é um índice baseado em 1 representando a posição de cada faixa na ordem de transmissão.

7. Você deve criar uma função chamada `remixPlaylist` que aceita um array de playlists e o número máximo de ocorrências permitidas por artista. A função deve retornar a programação final de transmissão como um array de objetos `{ slot, trackId }`, chamando `flattenPlaylists`, `scoreTracks`, `dedupeTracks`, `enforceArtistQuota` e `buildSchedule` na ordem.

# --hints--

Você deve ter uma função chamada `flattenPlaylists`.

```js
assert.isFunction(flattenPlaylists);
```

Você deve retornar um array vazio de `flattenPlaylists` quando a entrada não for um array.

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

Cada faixa retornada por `flattenPlaylists` deve incluir um campo `source` que é um array contendo o índice da playlist e o índice da faixa.

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

Você deve ter uma função chamada `scoreTracks`.

```js
assert.isFunction(scoreTracks)
```

Cada faixa retornada por `scoreTracks` deve incluir um campo numérico `score`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

Você deve calcular `score` usando um BPM alvo de `120` e esta fórmula: `votes * 10 - Math.abs(bpm - 120)`.

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

Você deve ter uma função chamada `dedupeTracks`.

```js
assert.isFunction(dedupeTracks)
```

Quando existirem valores duplicados de `trackId`, `dedupeTracks` deve manter apenas a primeira ocorrência da faixa.

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

Você deve ter uma função chamada `enforceArtistQuota`.

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` deve garantir que nenhum artista apareça mais do que `maxPerArtist` vezes removendo faixas extras enquanto mantém as mais antigas.

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

Você deve ter uma função chamada `buildSchedule`.

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` deve retornar um array de objetos com a forma `{ slot, trackId }`, onde `slot` começa em `1`.

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

Você deve ter uma função chamada `remixPlaylist`.

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` deve chamar as funções auxiliares na ordem para produzir a programação final.

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
