---
id: 69f0a37c3bf194a7c9c62492
title: Ein Verleihbuch für Geräte erstellen
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Sie sind Entwickler und sollen eine Ledger-Anwendung erstellen, um IT-Hardwaregeräte zu verwalten, die vom Service Desk Ihres Unternehmens bereitgestellt werden. Das Ledger ist ein Objekt, dessen Schlüssel Asset-Tags sind und dessen Werte Objekte sind:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten das Ledger als Objekt modellieren, bei dem jedes Gerät durch sein Asset-Tag indiziert ist und die Eigenschaften `type`, `status`, `borrower` (mit `name` und `email`) sowie `dueDate` besitzt.

1. Sie sollten eine `checkoutDevice(ledger, assetTag, borrower)`-Funktion implementieren, die ein Ledger-Objekt entgegennimmt und es klont, ohne das Original zu verändern. Im Klon soll sie die `name`- und `email`-Eigenschaften des `borrower`-Objekts setzen, den `status` auf `"CheckedOut"` aktualisieren und ein Objekt der Form `{ ledger: updatedLedger, message: confirmationString }` zurückgeben, wobei die `message` das Asset-Tag und den `name` des `borrower`-Objekts enthält.

   - Wenn das `assetTag` im Ledger nicht existiert, soll die Funktion das Ledger unverändert zurückgeben und eine `message` enthalten, die das Asset-Tag nennt und erklärt, dass es nicht gefunden wurde.
   - Wenn das Gerät bereits `"CheckedOut"` ist, soll die Funktion das Ledger unverändert zurückgeben, den aktuellen Borrower beibehalten und eine `message` enthalten, die das Asset-Tag nennt und erklärt, dass das Gerät bereits ausgeliehen ist.

1. Sie sollten eine `checkinDevice(ledger, assetTag)`-Funktion implementieren, die ein Ledger-Objekt entgegennimmt und es klont, ohne das Original zu verändern. Im Klon soll sie die `name`- und `email`-Eigenschaften des `borrower`-Objekts löschen, `dueDate` auf `""` zurücksetzen, den `status` auf `"CheckedIn"` aktualisieren und ein Objekt der Form `{ ledger: updatedLedger, message: confirmationString }` zurückgeben, wobei die `message` das Asset-Tag enthält.

   - Wenn das `assetTag` im Ledger nicht existiert, soll die Funktion das Ledger unverändert zurückgeben und eine `message` enthalten, die das Asset-Tag nennt und erklärt, dass es nicht gefunden wurde.

1. Sie sollten eine `listOverdueDevices(ledger, today)`-Funktion implementieren, die ein Ledger-Objekt entgegennimmt und ein Array der Geräteobjekte zurückgibt, die überfällig sind, sortiert nach `dueDate` in aufsteigender Reihenfolge. Ein Gerät ist überfällig, wenn sein `status` `"CheckedOut"` ist und sein `dueDate` strikt vor `today` liegt. Sowohl `today` als auch `dueDate` sind Strings im Monat/Tag/Jahr-Format, wobei Monat und Tag nullgefüllt sein können oder nicht (zum Beispiel sind `"9/5/2025"` und `"09/05/2025"` beide gültig).

1. Sie dürfen das JavaScript-`Date`-Objekt in Ihrem Code nirgendwo verwenden.

1. Sie sollten eine `serializeLedger(ledger)`-Funktion implementieren, die das Ledger-Objekt in einen JSON-String umwandelt, und eine `loadLedger(json)`-Funktion, die einen JSON-String zurück in ein JavaScript-Objekt konvertiert.

# --hints--

Sie sollten eine Funktion namens `checkoutDevice` erstellen.

```js
assert.isFunction(checkoutDevice);
```

Sie sollten eine Funktion namens `checkinDevice` erstellen.

```js
assert.isFunction(checkinDevice);
```

Sie sollten eine Funktion namens `listOverdueDevices` erstellen.

```js
assert.isFunction(listOverdueDevices);
```

Sie sollten eine Funktion namens `serializeLedger` erstellen.

```js
assert.isFunction(serializeLedger);
```

Sie sollten eine Funktion namens `loadLedger` erstellen.

```js
assert.isFunction(loadLedger);
```

Ihre `checkoutDevice`-Funktion darf das ursprüngliche Ledger nicht verändern.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

Ihre `checkoutDevice`-Funktion soll den `name` und die `email` des Borrowers setzen und den `status` auf `"CheckedOut"` setzen.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

Ihre `checkoutDevice`-Funktion soll eine Bestätigungsnachricht zurückgeben, die das Asset-Tag und den `name` des Borrowers enthält.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

Ihre `checkoutDevice`-Funktion soll nicht verwandte Geräte und unveränderte Felder im zurückgegebenen Ledger beibehalten.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Ihre `checkoutDevice`-Funktion soll das Ledger unverändert zurückgeben und eine Nachricht enthalten, die das Asset-Tag nennt, wenn das Asset-Tag nicht gefunden wird.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Ihre `checkoutDevice`-Funktion soll das Ledger unverändert zurückgeben, den aktuellen Borrower beibehalten und eine Nachricht enthalten, die das Asset-Tag nennt, wenn das Gerät bereits ausgeliehen ist.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

Ihre `checkinDevice`-Funktion darf das ursprüngliche Ledger nicht verändern.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

Ihre `checkinDevice`-Funktion soll den `name` und die `email` des Borrowers löschen, `dueDate` auf `""` zurücksetzen und den `status` auf `"CheckedIn"` setzen.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

Ihre `checkinDevice`-Funktion soll eine Bestätigungsnachricht zurückgeben, die das Asset-Tag enthält.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

Ihre `checkinDevice`-Funktion soll nicht verwandte Geräte und unveränderte Felder im zurückgegebenen Ledger beibehalten.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Ihre `checkinDevice`-Funktion soll das Ledger unverändert zurückgeben und eine Nachricht enthalten, die das Asset-Tag nennt, wenn das Asset-Tag nicht gefunden wird.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Ihre `listOverdueDevices`-Funktion soll ein Array zurückgeben.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

Ihre `listOverdueDevices`-Funktion soll nur überfällige Geräte mit dem `status` `"CheckedOut"` zurückgeben.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

Ihre `listOverdueDevices`-Funktion soll die überfälligen Geräte nach `dueDate` in aufsteigender Reihenfolge sortiert zurückgeben.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

Ihre `listOverdueDevices`-Funktion soll auch dann korrekt sortieren, wenn Monat oder Tag nicht nullgefüllt sind.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

Sie dürfen das JavaScript-`Date`-Objekt in Ihrem Code nirgendwo verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

Ihre `serializeLedger`-Funktion soll einen JSON-String zurückgeben, der das übergebene Ledger darstellt.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

Ihre `loadLedger`-Funktion soll das JavaScript-Objekt zurückgeben, das der JSON-String darstellt.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
