---
id: 66c057041df6394ca796bf33
title: 文作成メーカーを作ろう
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

このラボでは、文のテンプレートを使って2つの異なるストーリーを作成します。変数を使ってストーリーの異なる部分をストアし、その後ストーリーをコンソールに出力します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. 次の変数を`let`を使って宣言してください。
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. 上記の変数に任意の文字列の値を代入してください。

1. `firstStory`変数を宣言してください。
   
1. 次のストーリーテンプレートを使って最初のストーリーを作成し、`firstStory`変数に代入してください: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. 最初のストーリーを`"First story: [firstStory]"`というメッセージを使ってコンソールに出力してください。

1. `adjective`、`noun`、`verb`、`place`、`adjective2`、`noun2`の変数に新しい値を代入してください。

1. `secondStory`変数を宣言してください。

1. 同じテンプレートを使って別のストーリーを作成し、`secondStory`変数に代入してください。

1. 2つ目のストーリーを`"Second story: [secondStory]"`というメッセージを使ってコンソールに出力してください。


# --hints--

`adjective`変数を宣言してください。

```js
assert.isNotNull(adjective);
```

`noun`変数を宣言してください。

```js
assert.isNotNull(noun);
```

`verb`変数を宣言してください。

```js
assert.isNotNull(verb);
```

`place`変数を宣言してください。

```js
assert.isNotNull(place);
```

`adjective2`変数を宣言してください。

```js
assert.isNotNull(adjective2);
```

`noun2`変数を宣言してください。

```js
assert.isNotNull(noun2);
```

`adjective`変数に文字列の値を代入してください。

```js
assert.isString(adjective);
```

`noun`変数に文字列の値を代入してください。

```js
assert.isString(noun);
```

`verb`変数に文字列の値を代入してください。

```js
assert.isString(verb);
```

`place`変数に文字列の値を代入してください。

```js
assert.isString(place);
```

`adjective2`変数に文字列の値を代入してください。

```js
assert.isString(adjective2);
```

`noun2`変数に文字列の値を代入してください。

```js
assert.isString(noun2);
```

`firstStory`変数を宣言してください。

```js
assert.isNotNull(firstStory);
```

最初のストーリーには正しいストーリーフォーマットを使ってください: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。スペースに注意してください。

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

宣言した変数を正しい順序で使って最初のストーリーを組み立ててください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

`"First story: [firstStory]"`というメッセージを使って最初のストーリーをログに出力してください。

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`secondStory`変数を宣言してください。

```js
assert.isNotNull(secondStory);
```

2つ目のストーリー用に`adjective`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

2つ目のストーリー用に`noun`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

2つ目のストーリー用に`verb`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

2つ目のストーリー用に`place`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

2つ目のストーリー用に`adjective2`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

2つ目のストーリー用に`noun2`変数に再代入してください。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

2つ目のストーリーには正しいストーリーフォーマットを使ってください: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`。スペースに注意してください。

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

宣言した変数を正しい順序で使って2つ目のストーリーを組み立ててください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

`"Second story: [secondStory]"`というフォーマットを使って2つ目のストーリーをログに出力してください。

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory`は`secondStory`と同じであってはいけません。

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
