---
id: 69de4308e652aaeb5ccdb03e
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將練習使用 Python 的字串方法。

首先創建一個名為 `email` 的變數，並將字串 `alice.johnson@company.com` 指定給它。

# --hints--

你應該創建一個名為 `email` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

你應該將字串 `alice.johnson@company.com` 指定給 `email`。記得將文字用單引號或雙引號括起來。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
