---
id: 69e20876baed33de636b989a
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

有時你可能想要找出字串中特定字元的位置。Python 提供內建的 `.find()` 方法來達成這個目的。它會傳回一個整數，代表指定子字串第一次出現的位置（索引）。

這裡有一個範例：

```py
text = 'hello'
print(text.find('e'))  # 1
```

現在，對 `email` 呼叫 `.find()` 方法，並傳入 `'@'`，以找出 `email` 字串中 at 符號的位置。將結果指定給名為 `at_position` 的變數。

接著列印 `at_position`，讓輸出顯示為：`Position of @: 13`。

# --hints--

你應該創建一個名為 `at_position` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

你應該將 `at_position` 指定為 `email.find('@')` 的值。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

你應該列印 `at_position`，讓輸出顯示為 `Position of @: 13`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
