---
id: 69e20a0451c7a3ef189514d3
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

既然你已經知道 `@` 在電子郵件中的位置，就可以擷取放在 `@` 符號前面的使用者名稱。

在之前的工作坊中，你使用切片擷取字串的一部分。提醒你，語法是：

```python
text = 'hello world'
print(text[:5])  # hello
```

對 `email` 使用切片，擷取 `at_position` 之前的所有內容，並將結果存放在名為 `username` 的變數中。

然後列印 `username`，使輸出顯示為 `Username: alice.johnson`。

# --hints--

你應該創建一個名為 `username` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

你應該對 `email` 使用帶有 `at_position` 的切片，為 `username` 指定值。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

你應該列印 `username`，使輸出顯示為 `Username: alice.johnson`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
