---
id: 69e62ce96df237884da25282
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

`.lower()` 方法會將字串中的所有字元轉成小寫。這對於標準化電子郵件位址非常有用：

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

你可以用這種方式同時呼叫字串的 `.strip()` 和 `.lower()`：

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

創建一個變數 `messy_email`，並指定字串 `SARAH.DAVIS@COMPANY.COM`，字串前後各有兩個空白。

接著，對 `messy_email` 呼叫 `.strip().lower()`，並將結果指定給變數 `clean_email`。

然後列印兩者，使輸出顯示為：

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

你應該創建一個名為 `messy_email` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

你應該將 <code>'  SARAH.DAVIS@COMPANY.COM  '</code> 指定給 `messy_email`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

你應該創建一個名為 `clean_email` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

你應該對 `messy_email` 呼叫 `.strip().lower()`，並將結果指定給 `clean_email` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

你應該列印 `messy_email`，使輸出顯示為 `Original email: '  SARAH.DAVIS@COMPANY.COM  '`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你應該列印 `clean_email`，使輸出顯示為 `Cleaned email: 'sarah.davis@company.com'`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
