---
id: 69e6299710319e7cb8ebdc37
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

真實世界的資料通常很混亂，會有多餘的空格和不一致的大小寫。

創建一個名為 `messy_name` 的變數，並指定字串 `sARaH dAVis`，前後各有兩個空格。然後列印它，使輸出顯示為：

```md
Original name: '  sARaH dAVis  '
```

你將在接下來的步驟中清理這些資料。

# --hints--

你應該創建一個名為 `messy_name` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

你應該將 <code>'  sARaH dAVis  '</code> 指定給 `messy_name`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

你應該列印 `messy_name`，使輸出顯示為 `Original name: '  sARaH dAVis  '`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
