---
id: 69e6284f3b7659c5a05c2a79
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

你可以使用 `in` 關鍵字來檢查字串中是否存在子字串。如果子字串存在，會傳回 `True`，否則傳回 `False`。以下是一個範例：

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

創建一個名為 `has_at` 的變數，用來檢查字元 `@` 是否存在於 `email` 中。

然後列印 `has_at`，使輸出顯示為：`Has @ symbol: True`。

# --hints--

你應該創建一個名為 `has_at` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

你的 `has_at` 變數應該有值 `'@' in email`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

你應該列印 `has_at`，使輸出顯示為 `Has @ symbol: True`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
