---
id: 69e73a141addc6f442924b0d
title: 第 12 步驟
challengeType: 20
dashedName: step-12
---

# --description--

你已經在 `email` 變數上使用過 `.find()` 和切片。現在你將在 `clean_email` 上使用相同的技巧。

找出 `clean_email` 中 `@` 字元的位置並將其存入 `at_pos`。然後使用切片擷取 `@` 之前的所有內容並存入 `email_user`。

接著列印 `clean_email` 和 `email_user`，使輸出顯示為：

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

你應該創建一個名為 `at_pos` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

你應該在 `clean_email` 上呼叫 `.find('@')` 並將結果指定給 `at_pos` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

你應該創建一個名為 `email_user` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

你應該使用 `at_pos` 在 `clean_email` 上切片，並將結果指定給 `email_user` 變數。

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

你應該列印 `clean_email`，使輸出顯示為 `Email: sarah.davis@company.com`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你應該列印 `email_user`，使輸出顯示為 `Username: sarah.davis`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
