---
id: 69e75dc335b6ec9903a87d3e
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

你也可以使用 `.replace()` 來移除不需要的字元。例如，從電話號碼中移除連字號：

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

創建一個變數 `phone` 並指定字串 `123-456-7890`。然後將 `-` 替換成空字串，並將結果指定給名為 `phone_clean` 的變數。

接著列印兩者，使輸出顯示為：

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

你應該創建一個名為 `phone` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

你應該將字串 `123-456-7890` 指定給 `phone`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

你應該創建一個名為 `phone_clean` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

你應該對 `phone` 使用 `.replace('-', '')`，並將結果指定給 `phone_clean` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

你應該列印 `phone`，使輸出顯示為 `Phone with dashes: 123-456-7890`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

你應該列印 `phone_clean`，使輸出顯示為 `Phone without dashes: 1234567890`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
