---
id: 69e76e91402d7850510a5f18
title: 第 16 步
challengeType: 20
dashedName: step-16
---

# --description--

`.join(list)` 方法是 `.split()` 的相反。它使用分隔符將字串列表合併成一個字串。舉例來說：

```py
parts = ['apple', 'banana', 'cherry']
joined = ' - '.join(parts)
print(joined)  # Output: apple - banana - cherry
```

將 `address_parts` 的元素用兩側有空格的 `|` 字元連接起來，並將結果指定給名為 `rejoined` 的變數。

接著列印 `rejoined`，使輸出顯示為：`Rejoined: 123 Main Street | Springfield | IL`。

# --hints--

你應該創建一個名為 `rejoined` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('rejoined')
`) })
```

你應該將 `rejoined` 指定為 `' | '.join(address_parts)` 的值。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('rejoined').is_equivalent("rejoined = ' | '.join(address_parts)")
`) })
```

你應該列印 `rejoined`，使輸出顯示為 `Rejoined: 123 Main Street | Springfield | IL`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Rejoined: {rejoined}')",
    "print('Rejoined: ' + rejoined)",
    "print('Rejoined:', rejoined)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

--fcc-editable-region--

--fcc-editable-region--
```
