---
id: 69e73f4f5f986657647bbd63
title: 第 13 步驟
challengeType: 20
dashedName: step-13
---

# --description--

`.replace(old, new)` 方法會將字串中所有出現的子字串 `old` 替換成另一個字串 `new`。你也可以將它與其他方法串鏈使用。舉例來說：

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

將 `email_user` 中的 `.` 替換成一個空格，然後對結果呼叫 `.title()`。將結果指定給名為 `display_name` 的變數。

接著列印 `display_name`，讓輸出顯示為：`Display name: Sarah Davis`。

# --hints--

你應該創建一個名為 `display_name` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

你應該對 `email_user` 使用 `.replace('.', ' ')` 和 `.title()`，並將結果指定給 `display_name` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

你應該列印 `display_name`，讓輸出顯示為 `Display name: Sarah Davis`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
