---
id: 69e736c39dfa403496e8071e
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

`.upper()` 方法會將所有字元轉換成大寫。這對於名牌或標題非常有用：

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

創建一個名為 `name_badge` 的變數，並將對 `clean_name` 呼叫 `.upper()` 方法的結果指定給它。

然後列印 `name_badge`，使輸出顯示為：`Name badge: SARAH DAVIS`。

# --hints--

你應該創建一個名為 `name_badge` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

你應該對 `clean_name` 呼叫 `.upper()`，並將結果指定給 `name_badge` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

你應該列印 `name_badge`，使輸出顯示為 `Name badge: SARAH DAVIS`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
