---
id: 69e76d1997661be7f0040e2d
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

`.split(separator)` 方法會在每次遇到 `separator` 時，將字串拆分成子字串的列表。舉例來說：

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

創建一個變數 `full_address` 並指定字串 `123 Main Street, Springfield, IL` 給它。接著透過逗號加一個空白來拆分，創建 `address_parts`。

然後列印 `address_parts`，輸出結果應該是 `Address parts: ['123 Main Street', 'Springfield', 'IL']`。

# --hints--

你應該創建一個名為 `full_address` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

你應該將字串 `123 Main Street, Springfield, IL` 指定給 `full_address`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

你應該創建一個名為 `address_parts` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

你應該對 `full_address` 使用 `.split(', ')`，並將結果指定給 `address_parts` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

你應該列印 `address_parts`，輸出結果應該是 `Address parts: ['123 Main Street', 'Springfield', 'IL']`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
