---
id: 69e774b82fe0bfdb684f95b4
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

有時你需要檢查字串的開頭。Python 提供了 `.startswith()` 方法來做到這點。如果字串以指定的前置式開始，它會傳回 `True`；否則傳回 `False`。

這裡有一個範例：

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

創建一個變數 `employee_code` 並指定它為字串 `EMP-2024-SD`。接著創建一個 `starts_with_emp` 變數，用來檢查 `employee_code` 是否以 `EMP` 開頭。

# --hints--

你應該創建一個名為 `employee_code` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('employee_code')
`) })
```

你應該將 `EMP-2024-SD` 指定給 `employee_code`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('employee_code').is_equivalent("employee_code = 'EMP-2024-SD'")
`) })
```

你應該創建一個名為 `starts_with_emp` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('starts_with_emp')
`) })
```

你應該對 `employee_code` 使用 `.startswith('EMP')`，並將結果指定給 `starts_with_emp` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('starts_with_emp').is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

--fcc-editable-region--

--fcc-editable-region--
```
