---
id: 69eb3158c33d9a229a79951c
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

`.count(substring)` 方法會傳回子字串在字串中出現的次數。舉例來說：

```py
text = 'banana'
print(text.count('a'))  # Output: 3
```

創建兩個名為 `dot_count` 和 `dash_count` 的變數。`dot_count` 應該計算 `clean_email` 中出現了多少個 `.` 字元，而 `dash_count` 應該計算 `employee_code` 中出現了多少個 `-` 字元。

# --hints--

你應該創建一個名為 `dot_count` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dot_count')
`) })
```

你應該在 `clean_email` 上使用 `.count('.')` 來為 `dot_count` 指定值。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dot_count').is_equivalent("dot_count = clean_email.count('.')")
`) })
```

你應該創建一個名為 `dash_count` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dash_count')
`) })
```

你應該在 `employee_code` 上使用 `.count('-')` 來為 `dash_count` 指定值。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dash_count').is_equivalent("dash_count = employee_code.count('-')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Clean email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

--fcc-editable-region--

--fcc-editable-region--
```
