---
id: 6a71b3e49668c3be461257b5
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

`.title()` 方法會將每個字的第一個字母大寫。以下是一個範例：

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

創建一個名為 `clean_name` 的變數，並將對 `messy_name` 呼叫 `.title()` 的結果指定給它。

# --hints--

你應該創建一個名為 `clean_name` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

你應該對 `messy_name` 使用 `.title()`，並將結果指定給 `clean_name`。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
