---
id: 6a86be3e8cc29e9ae239e70a
title: 第 18 步
challengeType: 20
dashedName: step-18
---

# --description--

你可能也需要檢查字串的結尾。Python 提供 `.endswith()` 方法來做到這點。如果字串以指定的後綴結尾，該方法會傳回 `True`，否則傳回 `False`。

這裡有一個範例：

```python
code = 'EMP-001'
print(code.endswith('EMP')) # False
print(code.endswith('001')) # True
```

創建一個變數 `ends_with_com`，用來檢查 `clean_email` 是否以 `.com` 結尾。

# --hints--

你應該創建一個名為 `ends_with_com` 的變數。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('ends_with_com')
`) })
```

你應該對 `clean_email` 使用 `.endswith('.com')`，並將結果指定給 `ends_with_com` 變數。

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('ends_with_com').is_equivalent("ends_with_com = clean_email.endswith('.com')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')

--fcc-editable-region--

--fcc-editable-region--
```
