---
id: 6a870cc60c23f0332a2fe788
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

回想之前的工作坊，你可以使用 f-string 在字串中包含變數，只要在開頭引號前寫上 `f`，並將變數包在 `{}` 裡。

這裡提醒你它是如何運作的：

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

現在，列印 `email`，讓輸出顯示為：`Email: alice.johnson@company.com`。

# --hints--

你應該列印 `email`，讓輸出顯示為 `Email: alice.johnson@company.com`。

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
