---
id: 66ebd4ae2812430bb883c787
title: Erstellen eines Event-Hubs
challengeType: 25
dashedName: lab-event-hub
demoType: onClick
---

# --description--

In diesem Lab nutzen Sie die semantischen HTML-Elemente, um die Struktur einer Webseite zu erstellen. Sie fügen Inhalte und Bilder hinzu, damit sie wie ein echter Event-Hub aussieht.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein `header`-Element haben.
1. Innerhalb des `header`-Elements sollten Sie ein `h1`-Element haben, das den Text `Event Hub` enthält, sowie ein `nav`-Element.
1. Innerhalb des `nav`-Elements sollten Sie eine ungeordnete Liste mit zwei Elementen haben, die Links zu verschiedenen Abschnitten der Seite enthalten. Das erste Element sollte den Text `Upcoming Events` haben, das zweite den Text `Past Events`.
1. Jeder Link sollte durch ein `a`-Element mit einem `href`-Attribut dargestellt werden, das auf den entsprechenden Abschnitt der Seite verweist, `#upcoming-events` bzw. `#past-events`.
1. Sie sollten ein `main`-Element haben, das die verschiedenen Abschnitte der Seite enthält.
1. Innerhalb des `main`-Elements sollten Sie zwei `section`-Elemente haben.
1. Das erste `section`-Element sollte ein `id`-Attribut mit dem Wert `upcoming-events` haben.
1. Innerhalb des `#upcoming-events`-Abschnitts sollten Sie Folgendes haben:
   
   - Ein `h2`-Element mit dem Text `Upcoming Events`.
   - Zwei `article`-Elemente. Jeder Artikel sollte ein Event darstellen und folgendes enthalten:
      - Ein `h3`-Element für den Event-Titel.
      - Ein `p`-Element für die Event-Beschreibung. Sie können unten ein Datum hinzufügen, wenn Sie möchten.

1. Das zweite `section`-Element sollte ein `id`-Attribut mit dem Wert `past-events` haben.
1. Innerhalb des `#past-events`-Abschnitts sollten Sie Folgendes haben:

   - Ein `h2`-Element mit dem Text `Past Events`.
   - Zwei `article`-Elemente. Jedes Artikel-Element sollte ein vergangenes Event darstellen und folgendes enthalten:
     - Ein `h3`-Element für den Event-Titel,
     - Ein `p`-Element für die Event-Beschreibung. Sie können unten ein Datum hinzufügen, wenn Sie möchten.
     - Ein Bild-Element mit dem `src`-Attribut, das auf eine Bilddatei verweist, und dem `alt`-Attribut mit einer Beschreibung des Bildes.

**Hinweis:** Sie können beliebigen Text für die Event-Beschreibungen und Daten verwenden. Für die Bilder können Sie folgende Bild-URLs verwenden, wenn Sie möchten:

- `https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg`.
- `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`.

# --hints--

Sie sollten ein `header`-Element haben.

```js
assert.isNotNull(document.querySelector("header"));
```

Ihr `header`-Element sollte nach dem öffnenden `body`-Tag kommen.

```js
assert.equal(document.querySelector("body")?.firstElementChild?.tagName, "HEADER");
```

Innerhalb des `header`-Elements sollten Sie ein `h1`-Element haben, das den Text `Event Hub` enthält.

```js
const h1Element = document.querySelector('header h1');
assert.strictEqual(h1Element?.innerText.trim(), "Event Hub");
```

Innerhalb des `header`-Elements sollten Sie nach dem `h1`-Element ein `nav`-Element haben.

```js
assert.isNotNull(document.querySelector("header>h1+nav"));
```

Ihr `nav`-Element sollte eine ungeordnete Liste mit zwei Elementen enthalten.

```js
const liElements = document.querySelectorAll('header nav>ul>li');

assert.isNotNull('header nav>ul');
assert.strictEqual(liElements.length, 2);
```

Das erste Element in der ungeordneten Liste sollte ein Link sein.

```js
const firstLink = document.querySelectorAll('header nav ul li a')[0];
assert.exists(firstLink);
```

Das zweite Element in der ungeordneten Liste sollte ein Link sein.

```js
const secondLink = document.querySelectorAll('header nav ul li a')[1];
assert.exists(secondLink);
```

Der Text des ersten Elements in der ungeordneten Liste sollte `Upcoming Events` sein.

```js
const firstLink = document.querySelectorAll('header nav>ul>li>a')[0];
assert.strictEqual(firstLink.innerText.trim(), "Upcoming Events");
```

Das erste Element in der ungeordneten Liste sollte das `href`-Attribut mit dem Wert `#upcoming-events` haben.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[0];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#upcoming-events");
```

Der Text des zweiten Elements in der ungeordneten Liste sollte `Past Events` sein.

```js
const secondLink = document.querySelectorAll('header nav>ul>li>a')[1];
assert.strictEqual(secondLink.innerText.trim(), "Past Events");
```

Das zweite Element in der ungeordneten Liste sollte das `href`-Attribut mit dem Wert `#past-events` haben.

```js
const anchorElement = document.querySelectorAll("header nav>ul>li>a")[1];
const hrefAttribute = anchorElement?.getAttribute("href");
assert.strictEqual(hrefAttribute, "#past-events");
```

Sie sollten ein `main`-Element nach dem schließenden `header`-Element haben.

```js
const mainElement = document.querySelector("body>header+main");
assert.isNotNull(mainElement);
```

Innerhalb des `main`-Elements sollten Sie zwei `section`-Elemente haben.

```js
const sectionElements = document.querySelectorAll('body>header+main>section');
assert.strictEqual(sectionElements.length, 2);
```

Ihr erstes `section`-Element sollte ein `id`-Attribut mit dem Wert `upcoming-events` haben.

```js
const firstSection = document.querySelectorAll('body>header+main>section')[0];
const idAttribute = firstSection?.getAttribute("id");
assert.strictEqual(idAttribute, "upcoming-events");
```

Ihr zweites `section`-Element sollte ein `id`-Attribut mit dem Wert `past-events` haben.

```js
const secondSection = document.querySelectorAll('body>header+main>section')[1];
const idAttribute = secondSection?.getAttribute("id");
assert.strictEqual(idAttribute, "past-events");
```

Innerhalb des `#upcoming-events`-Abschnitts sollten Sie ein `h2`-Element mit dem Text `Upcoming Events` haben.

```js
const h2Element = document.querySelector('#upcoming-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Upcoming Events");
```

Innerhalb des `#upcoming-events`-Abschnitts sollten Sie zwei `article`-Elemente unterhalb des `h2`-Elements haben.

```js
const articleElements = document.querySelectorAll('#upcoming-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Beide `article`-Elemente innerhalb des `#upcoming-events`-Abschnitts sollten ein `h3`-Element für den Event-Titel haben.

```js
const h3Elements = document.querySelectorAll('#upcoming-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Beide `article`-Elemente innerhalb des `#upcoming-events`-Abschnitts sollten ein Absatz-Element für die Event-Beschreibung haben.

```js
const articles = document.querySelectorAll('#upcoming-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Innerhalb des `#past-events`-Abschnitts sollten Sie ein `h2`-Element mit dem Text `Past Events` haben.

```js
const h2Element = document.querySelector('#past-events h2');
assert.strictEqual(h2Element?.innerText.trim(), "Past Events");
```

Innerhalb des `#past-events`-Abschnitts sollten Sie zwei `article`-Elemente unterhalb des `h2`-Elements haben.

```js
const articleElements = document.querySelectorAll('#past-events h2 ~ article');
assert.strictEqual(articleElements.length, 2);
```

Beide `article`-Elemente innerhalb des `#past-events`-Abschnitts sollten ein `h3`-Element für den Event-Titel haben.

```js
const h3Elements = document.querySelectorAll('#past-events article h3');
assert.strictEqual(h3Elements.length, 2);
```

Beide `article`-Elemente innerhalb des `#past-events`-Abschnitts sollten ein Absatz-Element für die Event-Beschreibung haben.

```js
const articles = document.querySelectorAll('#past-events article');
assert.isNotEmpty(articles);
articles.forEach(article => {
    assert.isAtLeast(article.querySelectorAll('h3 ~ p').length, 1);
});
```

Beide `article`-Elemente innerhalb des `#past-events`-Abschnitts sollten ein Bild-Element haben.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);
```

Beide Bild-Elemente innerhalb des `#past-events`-Abschnitts sollten das `src`-Attribut haben, das auf eine Bilddatei verweist.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("src"));
}
```

Beide Bild-Elemente innerhalb des `#past-events`-Abschnitts sollten das `alt`-Attribut mit einer Beschreibung des Bildes haben.

```js
const imgElements = document.querySelectorAll('#past-events article img');
assert.strictEqual(imgElements.length, 2);

for (let img of imgElements) {
    assert.exists(img.getAttribute("alt"));
}
```

Jedes `h3`-Element sollte den Event-Titel haben.

```js
const eventTitles = document.querySelectorAll('h3');
assert.isNotEmpty(eventTitles);
eventTitles.forEach((eventTitle => assert.isNotEmpty(eventTitle.innerText.trim())));
```

Jedes `p`-Element sollte die Event-Beschreibung haben.

```js
const eventDescriptions = document.querySelectorAll('p');
assert.isNotEmpty(eventDescriptions);
eventDescriptions.forEach((eventDescription => assert.isNotEmpty(eventDescription.innerText.trim())));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Event Hub</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Event Hub</title>
</head>
<body>
    <header>
        <h1>Event Hub</h1>
        <nav>
            <ul>
                <li><a href="#upcoming-events">Upcoming Events</a></li>
                <li><a href="#past-events">Past Events</a></li>
            </ul>
        </nav>
    </header>
    <main>
        <section id="upcoming-events">
            <h2>Upcoming Events</h2>
            <article>
                <h3>AI & Machine Learning Conference 2024</h3>
                <p>Join us for a deep dive into the latest advancements in artificial intelligence and machine learning. Industry leaders will share insights and case studies on how AI is transforming various sectors.</p>
                <p>Date: August 10, 2024</p>
            </article>
            <article>
                <h3>Web Development Bootcamp</h3>
                <p>A hands-on workshop designed for developers looking to enhance their skills in modern web technologies including React, Node.js, and GraphQL. Perfect for both beginners and experienced developers.</p>
                <p>Date: September 5, 2024</p>
            </article>
        </section>
        <section id="past-events">
            <h2>Past Events</h2>
            <article>
                <h3>Cybersecurity Summit 2024</h3>
                <p>An event focusing on the latest trends and threats in cybersecurity. Experts discussed strategies for protecting data and ensuring privacy in an increasingly digital world.</p>
                <p>Date: June 15, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event1.jpg" alt="Image from Cybersecurity Summit 2024">
            </article>
            <article>
                <h3>Blockchain Expo 2024</h3>
                <p>A comprehensive event covering the future of blockchain technology. Topics included decentralized finance (DeFi), smart contracts, and the impact of blockchain on various industries.</p>
                <p>Date: July 20, 2024</p>
                <img src="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg" alt="Image from Blockchain Expo 2024">
            </article>
        </section>
    </main>
</body>
</html>
```

