---
id: 686ccc2c8b967e17ab18d593
title: Tekeleza njia ya bisection
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

Njia ya bisection, inayojulikana pia kama njia ya utafutaji wa binary, hutumia utafutaji wa binary kupata mizizi ya kitendakazi chenye thamani halisi. Inafanya kazi kwa kupunguza eneo ambalo mzizi wa mraba uko ndani yake hadi itakapokaribia thamani ndani ya uvumilivu uliowekwa.

Kwa mfano, ikiwa uvumilivu ni `0.01`, njia ya bisection itaendelea kugawanya eneo hilo nusu hadi tofauti kati ya mipaka ya juu na ya chini iwe ndogo au sawa na `0.01`.

Katika maabara hii, utatekeleza kitendakazi kinachotumia njia ya bisection kupata mzizi wa mraba wa nambari.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kufafanua kitendakazi kinachoitwa `square_root_bisection` chenye vigezo vitatu:
    - Nambari unayotaka kupata mzizi wake wa mraba.
    - Uvumilivu ambao ni ukingo wa makosa unaokubalika kwa matokeo. Unapaswa kuweka thamani ya uvumilivu kama chaguo-msingi.
    - Idadi ya juu ya mizunguko ya kufanya. Unapaswa kuweka idadi ya mizunguko kama chaguo-msingi.

1. Kitendakazi cha `square_root_bisection` kinapaswa:
    - Kusababisha `ValueError` yenye ujumbe `Square root of negative number is not defined in real numbers` ikiwa nambari iliyopitishwa ni hasi.
    - Kwa nambari `0` na `1`, chapisha ujumbe: `The square root of [number] is [number]` na rudisha nambari yenyewe kama mzizi wa mraba.
    - Kwa nambari nyingine yoyote chanya, chapisha mzizi wa mraba wa takriban kwa ujumbe: `The square root of [square_target] is approximately [root]` na rudisha thamani ya mzizi iliyohesabiwa.
    - Ikiwa hakuna thamani inayokidhi sharti la uvumilivu, chapisha ujumbe wa kushindwa: `Failed to converge within [maximum] iterations` na rudisha `None`.

**Kumbuka**: Huwezi import moduli yoyote kwa maabara hii.

# --hints--

Huwezi import moduli yoyote.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

Unapaswa kuwa na kitendakazi kinachoitwa `square_root_bisection`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

Kitendakazi chako cha `square_root_bisection` kinapaswa kuwa na vigezo vitatu.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

Unapaswa kuweka thamani chaguo-msingi kwa uvumilivu na idadi ya juu ya mizunguko.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

Kitendakazi chako cha `square_root_bisection` kinapaswa kusababisha `ValueError` yenye ujumbe `Square root of negative number is not defined in real numbers` wakati nambari iliyopitishwa ni hasi.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` inapaswa kurudisha `0`.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` inapaswa kuchapisha `The square root of 0 is 0`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` inapaswa kurudisha nambari kati ya `0.03162267660168379` na `0.031622876601683794`.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` inapaswa kuchapisha `The square root of 0.001 is approximately X`, ambapo `X` ni nambari kati ya `0.03162267660168379` na `0.031622876601683794`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` inapaswa kurudisha nambari kati ya `0.4999999` na `0.5000001`.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` inapaswa kuchapisha `The square root of 0.25 is approximately X`, ambapo `X` ni nambari kati ya `0.4999999` na `0.5000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` inapaswa kurudisha `1`.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` inapaswa kuchapisha `The square root of 1 is 1`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` inapaswa kurudisha nambari kati ya `8.999` na `9.001`.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` inapaswa kuchapisha `The square root of 81 is approximately X`, ambapo `X` ni nambari kati ya `8.999` na `9.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` inapaswa kurudisha nambari kati ya `14.999` na `15.001`.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` inapaswa kuchapisha `The square root of 225 is approximately X`, ambapo `X` ni nambari kati ya `14.999` na `15.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` inapaswa kurudisha nambari kati ya `14.99999` na `15.00001`.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` inapaswa kuchapisha `The square root of 225 is approximately X`, ambapo `X` ni nambari kati ya `14.99999` na `15.00001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` inapaswa kurudisha nambari kati ya `14.9999999` na `15.0000001`.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` inapaswa kuchapisha `The square root of 225 is approximately X`, ambapo `X` ni nambari kati ya `14.9999999` na `15.0000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` inapaswa kurudisha `None`.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` inapaswa kuchapisha `Failed to converge within 10 iterations`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
