---
id: 66bb6a9c2dd58b73cd759034
title: 성적부 앱 만들기
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `getAverage`라는 이름의 함수가 있어야 하며, 이 함수는 시험 점수 배열을 매개변수로 받아 평균 점수를 반환해야 합니다.
1. `getGrade`라는 이름의 함수가 있어야 하며, 이 함수는 학생 점수를 매개변수로 받아 점수에 따른 학점 문자열을 반환해야 합니다. 점수와 해당 학점은 다음과 같습니다:

    | 점수 범위 | 학점   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. `hasPassingGrade`라는 이름의 함수가 있어야 하며, 이 함수는 점수를 매개변수로 받아 점수가 합격 학점에 해당하면 `true`를, 그렇지 않으면 `false`를 반환해야 합니다.
1. `hasPassingGrade` 함수는 `getGrade` 함수를 사용해 학점을 얻고, 그 학점이 합격인지 판단해야 합니다. 합격 학점은 `"F"`가 아닌 모든 학점입니다.
1. `studentMsg`라는 이름의 함수가 있어야 하며, 이 함수는 점수 배열과 학생 점수를 매개변수로 받습니다. 함수는 다음 형식의 문자열을 반환해야 합니다:

    - 학생이 과목을 통과했으면 `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`,
    - 학생이 과목을 실패했으면 `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`.

    `average-goes-here`은 전체 점수의 평균으로, `grade-goes-here`는 학생의 학점으로 바꾸세요.
    평균 점수는 `getAverage`를, 학생 학점은 `getGrade`를 사용해 구하세요.

# --hints--

`getAverage`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(getAverage);
```

`getAverage` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])`는 `71.7`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])`는 `85.4`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])`는 `92.4`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])`는 `47.25`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])`는 `47.125`를 반환해야 합니다.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

`getAverage` 함수는 시험 점수들의 평균 점수를 반환해야 합니다.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

`getGrade`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(getGrade);
```

`getGrade` 함수는 점수가 `"A+"`일 때 `100`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(100), "A+");
```

`getGrade` 함수는 점수가 `"A"`과 `90` 사이일 때 `99`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

`getGrade` 함수는 점수가 `"B"`과 `80` 사이일 때 `89`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

`getGrade` 함수는 점수가 `"C"`와 `70` 사이일 때 `79`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

`getGrade` 함수는 점수가 `"D"`과 `60` 사이일 때 `69`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

`getGrade` 함수는 점수가 `"F"`과 `0` 사이일 때 `59`를 반환해야 합니다.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


`hasPassingGrade`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(hasPassingGrade);
```

`hasPassingGrade` 함수는 불리언 값을 반환해야 합니다.

```js
assert.isBoolean(hasPassingGrade(100));
```

`hasPassingGrade` 함수는 학점이 `true`일 때 `"A"`를 반환해야 합니다.

```js
assert.isTrue(hasPassingGrade(100));
```

`hasPassingGrade` 함수는 학점이 `false`일 때 `"F"`를 반환해야 합니다.

```js
assert.isFalse(hasPassingGrade(53));
```

`hasPassingGrade` 함수는 모든 합격 학점에 대해 `true`를 반환해야 합니다.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

`studentMsg`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)`는 다음 메시지를 반환해야 합니다: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)`는 다음 메시지를 반환해야 합니다: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)`는 다음 메시지를 반환해야 합니다: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)`는 다음 메시지를 반환해야 합니다: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

`studentMsg` 함수는 학생 점수와 수업 평균에 따라 올바른 메시지를 반환해야 합니다.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
