---
id: 6823d6244511f252c8300eed
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Cuando necesite agregar un párrafo a una página web, puede usar el elemento `p` de esta manera:

```html
<p>This is a paragraph element.</p>
```

Convierta `Learn the skills to become a full stack developer` en un elemento de párrafo.

# --hints--

Su elemento `p` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

Su elemento `p` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen esta sintaxis: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

El texto de su elemento `p` debe ser `Learn the skills to become a full stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full\s*stack\s*developer\s*\<\/p\s*\>/i);
```

Su elemento `p` debe estar debajo de su elemento `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full stack developer
--fcc-editable-region--
```
