---
id: 6823ac607bfdbc46331b2559
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML significa HyperText Markup Language. Es el código que define la estructura y el contenido de una página web. Este es tu editor de código, donde escribirás HTML.

Busca la línea 1 en el editor y escribe este texto:

```md
Welcome to freeCodeCamp
```

Cuando termines, haz clic en el botón "verificar el código" para ver si es correcto.

# --hints--

Debes tener el texto `Welcome to freeCodeCamp` en tu editor. Revista la ortografía cuidadosamente.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
