---
id: 6823e036ea4b71553558c01b
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Observe que cada encabezado se ve un poco diferente en la vista previa. Lo mismo ocurre con el párrafo.

A continuación, añada otro elemento `p` que muestre el siguiente texto:

```md
HTML represents the content and structure of a webpage
```

# --hints--

Debería tener un segundo elemento `p`.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Su elemento `p` debería tener el texto `HTML represents the content and structure of a webpage`. Puede copiar el texto de las instrucciones.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
