---
id: 6a89c69c91bc7365fbd4b733
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

為你的 `countup` 函式（程式）添加基本案例。將 `pass` 替換成一個 `if` 述語，檢查 `number` 是否小於 `1`。在這個 `if` 述語內，傳回一個空的列表。

# --hints--

當 `number` 小於 `1` 時，你的 `countup` 函式（程式）應該傳回一個空的列表。

```js
({ test: () => runPython(`assert countup(0) == []`) })
```

當 `number` 大於或等於 `1` 時，你的 `countup` 函式（程式）不應該傳回空的列表。

```js
({ test: () => runPython(`assert countup(1) is None`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
