---
id: 6a89c69c6af682856558d78b
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

遞迴呼叫會傳回一個包含較小數字的列表。使用 `append()` 將 `number` 添加到 `count_list`，然後傳回 `count_list`。

# --hints--

你應該將 `number` 附加到 `count_list`。

```js
({ test: () => runPython(`
_calls = _Node(_code).find_function("countup").find_calls("append")
assert len(_calls) == 1
assert _calls[0].is_equivalent("count_list.append(number)")
`) })
```

你的 `countup` 函式應該從中傳回 `count_list`。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_return("count_list")`) })
```

當以 `3` 呼叫時，你的 `countup` 函式應該傳回 `[1, 2, 3]`。

```js
({ test: () => runPython(`assert countup(3) == [1, 2, 3]`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
--fcc-editable-region--
    
--fcc-editable-region--
```
