---
id: 6a89c69c8dbfda74151fa0b6
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建一個 `countup` 函式（程式），使用遞迴來生成從 `1` 到指定數字的數字列表。

首先定義一個名為 `countup` 的函式（程式），它有一個名為 `number` 的參數。為了讓函式在空的時候仍然有效，請在函式主體中加入 `pass` 關鍵字。

# --hints--

你應該有一個名為 `countup` 的函式（程式），並帶有一個名為 `number` 的參數。

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("countup")
assert not _func.is_empty()
assert _func.has_args("number")
`) })
```

你的 `countup` 函式（程式）應該包含 `pass` 關鍵字。

```js
({ test: () => runPython(`assert _Node(_code).find_function("countup").has_pass()`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
