---
id: 6a89c69c99298438d6fbe60b
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

現在加入遞迴情況。在 `if` 述語之後，創建一個名為 `count_list` 的變數。將呼叫 `countup` 並傳入 `number - 1` 的結果指定給它。

每次呼叫都使用較小的數字，因此遞迴會朝向基底情況前進。

# --hints--

你應該將 `countup(number - 1)` 指定給名為 `count_list` 的變數。

```js
({ test: () => runPython(`
_variable = _Node(_code).find_function("countup").find_variable("count_list")
assert _variable.is_equivalent("count_list = countup(number - 1)")
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
--fcc-editable-region--
    
--fcc-editable-region--
```
