---
id: 6a89c69c111694edc6d1c628
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

透過在 `print()` 呼叫中傳入 `5`，測試你的 `countup` 函式（程式）。輸出應該是 `[1, 2, 3, 4, 5]`。

這樣，countup 函式（程式）工作坊就完成了。

# --hints--

你應該在 `print()` 函式（程式）中呼叫 `countup(5)`。

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(countup(5))")
assert countup(5) == [1, 2, 3, 4, 5]
`) })
```

# --seed--

## --seed-contents--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
def countup(number):
    if number < 1:
        return []
    count_list = countup(number - 1)
    count_list.append(number)
    return count_list

print(countup(5))
```
