---
id: 0095d0567b6817114989a1dd
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将用 JavaScript 实现最短路径算法。你需要编写一个函数，计算加权图中节点之间的最短路径，并返回距离和所经过的路径。

例如，给定一个城市通过不同距离的道路连接的图，算法会找到从一个城市到另一个城市的最短路线。如果你想从城市 A 旅行到城市 D，算法可能会发现走 A ⇨ B ⇨ C ⇨ D（总计：15 公里）比直接走 A ⇨ D（20 公里）更短。

开始时，定义一个名为 `INF` 的变量，并赋值为 `Infinity`。稍后你会用它来表示两个节点之间没有直接连接。

# --hints--

你应该定义一个名为 `INF` 的变量。

```js
assert.isDefined(INF);
```

你的 `INF` 变量应该赋值为 `Infinity`。

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
