---
id: 69c918841987bf16c1d1bac2
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

在循环内部，你需要扫描每个节点，查找尚未访问且已知距离最小的节点。

在主循环内添加一个 `for` 循环，迭代每个节点。使用 `nodeNo` 作为循环变量，从 `0` 到 `n`。

暂时将体留空。

# --hints--

你应该在主循环内有第二个（内层）`for` 循环，使用 `nodeNo` 作为循环变量。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

--fcc-editable-region--
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;

  }
--fcc-editable-region--
}
```
