---
id: 69c918841987bf16c1d1baba
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

你需要一个二维数组来表现图的邻接矩阵。每个值代表两个节点之间边的权值（距离）。值为 `INF` 表示没有直接的边。

创建另一个名为 `adjMatrix` 的变量，并赋值为一个包含以下值的二维数组：

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

你应该创建一个名为 `adjMatrix` 的变量。

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` 应该是一个有 6 行 6 列的二维数组。

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` 应该包含正确的值。

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
