---
id: 69c918841987bf16c1d1bac4
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

如果条件为真，当前节点就是迄今为止找到的最佳未访问候选节点。你需要更新你的跟踪变量以反映这一点。

在 `if` 语句内，将 `minDistance` 更新为 `distances[nodeNo]`，并将 `current` 更新为 `nodeNo`。

# --hints--

你应该在 `minDistance` 块内将 `distances[nodeNo]` 设置为 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /minDistance\s*=\s*distances\s*\[\s*nodeNo\s*\]/
);
```

你应该在 `current` 块内将 `nodeNo` 设置为 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /current\s*=\s*nodeNo/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
