---
id: 69c918841987bf16c1d1babf
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

你还需要跟踪到达每个节点的实际路径。你将保存一个路径列表，其中每个条目都是一个节点索引数组，表现所经过的路线。

最初，每个节点的路径只包含它自己。你可以使用 `Array.from()` 来创建这个结构：

```js
Array.from({ length: n }, (_, i) => [i]);
```

这会创建一个包含 `n` 个元素的数组，其中索引为 `i` 的元素是 `[i]`。

仍然在 `shortestPath` 函数内，创建一个名为 `paths` 的变量，并赋值为 `Array.from({ length: n }, (_, i) => [i])`。

# --hints--

你应该在 `paths` 函数内创建一个变量 `shortestPath`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

你的 `paths` 变量应该使用 `Array.from` 来初始化。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

你的 `paths` 变量应该使用 `{ length: n }` 的第一个参数 `Array.from`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

`Array.from` 的第二个参数应该是一个回调 `(_, i) => [i]`，其中 `_` 是未使用元素的占位符，`i` 是索引。它应该返回 `[i]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--

  --fcc-editable-region--
}
```
