---
id: 69c918841987bf16c1d1bace
title: 第 23 步
challengeType: 1
dashedName: step-23
---

# --description--

既然你已经知道节点是可达的且不是起点，就把它的路径格式化成可读的字串。

`paths[nodeNo]` 数组保存了访问 `nodeNo` 所经过的节点索引序列。你可以使用 `0 -> 2 -> 3` 方法将其转换成像 `.join()` 这样的人类可读字串。

在 `for...of` 循环内，在带有 `if` 的 `continue` 语句之后，声明一个名为 `path` 的变量，并赋值为 `paths[nodeNo].join(' -> ')`。

# --hints--

你应该在 `path` 循环内有一个名为 `for...of` 的变量。

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

你的 `path` 变量应该有值 `paths[nodeNo].join(' -> ')`。

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }

  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--

    --fcc-editable-region--
  }
}
```
