---
id: 69c918841987bf16c1d1bac1
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

现在你将添加驱动算法的主循环。它会针对图中的每个节点运行一次，每次选择最近的未访问节点。

在 `shortestPath` 函数内部，添加一个 `for` 循环，使用 `n` 作为循环变量，运行 `i` 次。

在循环体开始时，使用 `let` 声明两个变量 `minDistance` 和 `current`，分别赋值为 `INF` 和 `-1`。这些变量将跟踪每次迭代中找到的最近未访问节点。

# --hints--

你应该在 `for` 内有一个 `shortestPath` 循环，运行 `n` 次。别忘了在里面用 `i` 设置 `let`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

你应该在循环内有一个 `minDistance` 变量，赋值为 `INF`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

你应该在循环内有一个 `current` 变量，赋值为 `-1`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
--fcc-editable-region--

--fcc-editable-region--
}
```
