---
id: 69c918841987bf16c1d1bac8
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

在更新距离之前，你需要确认邻居节点值得考虑。必须存在一条实际的边（`distance !== INF`），并且该邻居节点尚未被访问。

在邻居 `for` 的循环内，添加一个 `if` 语句，用来检查 `distance` 是否不严格相等于 `INF` **且** `nodeNo` 尚未被访问。

在 `if` 块内，声明一个变量 `newDistance`，赋值为 `distances[current] + distance`。这是通过当前节点到达邻居的总成本。

# --hints--

你应该添加一个条件为 `if` 的 `distance !== INF && !visited[nodeNo]` 语句。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

你应该在 `newDistance` 块内声明一个值为 `distances[current] + distance` 的 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
