---
id: 69c918841987bf16c1d1bac9
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

现在你应该查看通过当前节点是否能为邻居提供更短的路径。如果 `newDistance` 比已经保存的更优，那么你应该更新它。

在现有的 `if` 块内，添加一个嵌套的 `if` 语句，检查 `newDistance` 是否小于 `distances[nodeNo]`。在那个嵌套的 `if` 块内，将 `distances[nodeNo]` 更新为 `newDistance`。

# --hints--

你应该添加一个条件为 `if` 的嵌套 `newDistance < distances[nodeNo]` 语句。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*newDistance\s*<\s*distances\s*\[\s*\w+\s*\]\s*\)/
);
```

你应该在嵌套的 `distances[nodeNo]` 块内将 `newDistance` 更新为 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*\w+\s*\]\s*=\s*newDistance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
