---
id: 69c918841987bf16c1d1bac7
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

现在当前节点已被标记为已访问，你需要查看它的所有邻居，看看是否可以通过当前节点更高效地到达它们。

在 `visited[current] = true` 之后，添加一个邻居 `for` 循环，迭代 `nodeNo` 到 `0` 之间的每个 `n`。注意这是与第一个循环同级的另一个内层循环。

在循环内部，声明一个名为 `distance` 的变量，并赋值为 `matrix[current][nodeNo]`。这将给出当前节点与其邻居之间的边的权值。

# --hints--

你应该在标记当前节点为已访问后添加一个邻居 `for` 循环。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

邻居 `for` 循环应迭代从 `0` 到 `n`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

你应该在邻居循环内有一个 `distance` 变量赋值为 `matrix[current][nodeNo]`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
