---
id: 69c918841987bf16c1d1bac3
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

在内层 `for` 循环中，你需要查看当前节点是否未被访问且比你目前找到的最佳节点更近。

添加一个 `if` 语句，用来检查 `nodeNo` 是否未被访问且 `distances[nodeNo]` 是否小于 `minDistance`。

暂时将体留空。

# --hints--

你应该在内层 `if` 循环中有一个条件为 `!visited[nodeNo] && distances[nodeNo] < minDistance` 的 `for` 语句。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*!visited\s*\[\s*nodeNo\s*\]\s*&&\s*distances\s*\[\s*nodeNo\s*\]\s*<\s*minDistance\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
