---
id: 69c918841987bf16c1d1babb
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

现在你将创建执行所有工作的主函数。

创建一个名为 `shortestPath` 的函数，它接受三个参数：`matrix`、`startNode` 和 `targetNode`。

在接下来的几个步骤中，你将构建这个函数的逻辑。

# --hints--

你应该创建一个名为 `shortestPath` 的函数。

```js
assert.isFunction(shortestPath);
```

你的 `shortestPath` 函数应该有参数 `matrix`、`startNode` 和 `targetNode`。

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
