---
id: 69c918841987bf16c1d1bac5
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在扫描所有节点以查找最近的未访问节点后，你需要处理没有找到有效节点的情况。当所有剩余节点都无法到达时，就会发生这种情况。

在内层 `for` 循环之后，但仍在外层循环内，添加一个 `if` 语句，检查 `current` 是否严格相等于 `-1`。如果为真，则在 `break` 块内使用 `if` 来提前停止外层循环。

# --hints--

你应该添加一个 `if` 语句来检查 `current === -1`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

你应该在 `break` 块内使用 `if`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
