---
id: 69c918841987bf16c1d1bac0
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

随着算法的运行，你需要跟踪哪些节点已经被处理过，以避免重复访问它们。

创建一个名为 `visited` 的变量，并赋值为 `new Array(n).fill(false)`。

# --hints--

你应该创建一个名为 `visited` 的变量。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

你的 `visited` 变量应该被赋值为 `new Array(n).fill(false)`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
