---
id: 69c918841987bf16c1d1babe
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

从起始节点到它自身的距离总是 `0`。你必须更新 `distances` 数组以反映这一点。

在 `distances` 数组之后，将 `distances[startNode]` 设置为 `0`。

# --hints--

你应该将 `distances[startNode]` 设置为 `0`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*startNode\s*\]\s*=\s*0/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);

}
--fcc-editable-region--
```
