---
id: 69c918841987bf16c1d1babc
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

你需要保存图中节点的数量。

在 `shortestPath` 函数内，创建一个名为 `n` 的变量，并赋值为 `matrix.length`。

# --hints--

你应该在 `n` 函数内创建一个变量 `shortestPath`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

你的 `n` 变量应当具有值 `matrix.length`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {

}
--fcc-editable-region--
```
