---
id: 69c918841987bf16c1d1babd
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

你需要跟踪从起始节点到每个其他节点的已知最短距离。开始时，你会假设每个节点的距离都是无限远。

在 JavaScript 中，你可以使用 `new Array(n).fill(value)` 创建一个预填充了某个值的数组。例如，`new Array(3).fill(0)` 会创建 `[0, 0, 0]`。

在 `shortestPath` 函数内部，创建一个名为 `distances` 的变量，并将 `new Array(n).fill(INF)` 赋值给它。

# --hints--

你应该在 `distances` 内创建一个名为 `shortestPath` 的变量。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

你的 `distances` 变量应该被赋值为 `new Array(n).fill(INF)`。

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
