---
id: 5cdafbb0291309899753167f
title: 创建一个 JavaScript Promise
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

JavaScript 中的 promise 就是字面意思——你用它来承诺做某事，通常是异步的。当任务完成时，你要么履行承诺，要么未能履行。`Promise` 是一个构造函数，所以你需要使用 `new` 关键字来新建一个。它接受一个函数作为参数，该函数有两个参数——`resolve` 和 `reject`。这两个方法用来决定 promise 的结果。语法如下：

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

创建一个名为 `makeServerRequest` 的 promise。 给构造器函数传入 `resolve` 和 `reject` 两个参数。

# --hints--

应该给名为 `makeServerRequest` 的变量指定一个 promise。

```js
assert(makeServerRequest instanceof Promise);
```

promise 应该接收一个函数做为参数，该函数应该包含 `resolve` 和 `reject` 两个参数。

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
