---
id: 67c562286b29447da020d407
title: 일회용 비밀번호 생성기 만들기
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

이 실습에서는 6자리 OTP(일회용 비밀번호)를 생성하여 사용자에게 표시합니다. OTP는 5초 후 만료되며, 사용자가 `"Generate New OTP"` 버튼을 클릭하면 새 OTP가 생성됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `useEffect` 훅을 사용하여 카운트다운 타이머를 관리해야 합니다.
2. `OTPGenerator` 컴포넌트는 클래스 이름이 `div`인 `container` 요소를 반환해야 합니다.
3. 클래스가 `div`인 `container`는 다음 요소들을 포함해야 합니다:

    - ID가 `h1`이고 텍스트가 `otp-title`인 `"OTP Generator"` 요소.
    - ID가 `h2`인 `otp-display` 요소로, 메시지 `"Click 'Generate OTP' to get a code"`를 표시하거나 생성된 OTP가 있으면 이를 보여야 합니다.
    - ID가 `p`이고 `otp-timer` 속성이 유효한 값으로 설정된 `aria-live` 요소로:
        - 처음에는 비어 있어야 합니다.
        - 버튼이 클릭된 후 `"Expires in: X seconds"`를 표시하며, 여기서 `X`은 OTP 만료 전 남은 시간을 나타냅니다.
        - 카운트다운이 `"OTP expired. Click the button to generate a new OTP."`에 도달하면 메시지 `0`를 보여야 합니다.
    - ID가 `button`이고 레이블이 `generate-otp-button`인 `"Generate OTP"` 요소. 클릭 시 새 OTP를 생성하고 5초 카운트다운을 시작해야 합니다.
    - 카운트다운이 진행 중일 때 `"Generate OTP"` 버튼은 비활성화되어야 합니다.

4. 불필요한 업데이트를 방지하기 위해 카운트다운 타이머가 `0`초에 도달하면 자동으로 정지되도록 해야 합니다.
5. 생성된 OTP는 6자리 숫자여야 합니다.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

기존 `const { useState, useEffect, useRef } = React;` 할당을 코드에서 제거하지 마세요.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

클래스가 `div`인 `.container`를 렌더링해야 합니다.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

`.container`는 ID가 `h1`이고 텍스트가 `otp-title`인 `"OTP Generator"` 요소를 포함해야 합니다.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

`.container`는 생성된 OTP를 표시하는 ID가 `h2`인 `otp-display` 요소를 포함해야 합니다.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

처음에는 `h2` 내부의 `.container`가 메시지 `"Click 'Generate OTP' to get a code"`를 표시해야 합니다.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

`.container` 요소는 ID가 `p`인 `otp-timer` 요소를 포함해야 합니다.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

`p` 요소는 `aria-live` 속성이 `assertive` 또는 `polite`로 설정되어야 합니다.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

처음에는 `p` 요소가 비어 있어야 합니다.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

`.container`는 ID가 `button`이고 텍스트가 `generate-otp-button`인 `"Generate OTP"` 요소를 포함해야 합니다.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

버튼이 클릭되면 ID가 `h2`인 `otp-display` 요소에 새 OTP가 표시되어야 합니다.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

생성된 OTP는 6자리 숫자여야 합니다.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

버튼 클릭 시 ID가 `p`인 `otp-timer` 요소에 5초 카운트다운이 표시되어야 합니다.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

ID가 `p`인 `otp-timer` 요소의 메시지는 매초 업데이트되어 남은 시간을 보여야 합니다.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

카운트다운 타이머가 실행 중일 때 `"Generate OTP"` 버튼은 비활성화되어야 합니다.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

카운트다운 타이머가 0에 도달하면 `"Generate OTP"` 버튼은 활성화되어야 합니다.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

카운트다운 타이머가 `0`에 도달하면 메시지 `OTP expired. Click the button to generate a new OTP.`를 표시해야 합니다.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

`OTPGenerator` 컴포넌트를 내보내야 합니다.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
