---
id: 68b94f23f7b3ee14078c89c5
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 `iframe` 요소를 사용해 비디오를 표시할 것입니다. 기본 HTML 보일러플레이트가 준비되어 있습니다.

먼저 `h1` 요소를 만들고 텍스트로 `iframe Video Display`를 입력하세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'))
```

`h1` 요소에는 텍스트 `iframe Video Display`가 있어야 합니다.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
