---
id: 68ba9c4f1688914d72876094
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

`iframe` 요소에 관한 첫 번째 수업에서 `img`처럼 교체된 요소라는 것을 배웠습니다. 이는 높이와 너비를 결정하기 위해 `width`와 `height` 속성도 사용할 수 있다는 뜻입니다.

`iframe` 요소에 `width`를 `560`로, `height`를 `315`로 지정하세요.

# --hints--

`iframe` 요소에 `width` 속성을 `560`로 설정해야 합니다.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

`iframe` 요소에 `height` 속성을 `315`로 설정해야 합니다.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
