---
id: 68ba9c4f1688914d72876096
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

속성 중 하나는 `allow`입니다. 이것은 브라우저가 `iframe`에서 사용할 수 있는 기능을 알려주는 권한 목록과 같습니다.

다음은 `iframe` 속성이 있는 `allow` 요소입니다:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

값이 `allow`, `accelerometer`, 그리고 `autoplay`인 `clipboard-write` 속성을 추가하세요.

`accelerometer`는 `iframe`가 기기 기울기와 회전 같은 동작 센서를 사용할 수 있게 합니다. `autoplay`는 비디오가 자동으로 재생되도록 허용하며, `clipboard-write`는 iframe이 사용자 클립보드에 데이터를 쓸 수 있게 합니다.

# --hints--

`iframe` 요소에 `allow` 속성을 `accelerometer`로 설정하세요.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

`allow` 요소의 `iframe` 속성에 `autoplay`가 값 중 하나로 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

`allow` 요소의 `iframe` 속성에 `clipboard-write`가 값 중 하나로 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
