---
id: 68ba9c4f1688914d72876097
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

기존 `encrypted-media` 속성 값에 `gyroscope`, `web-share`, 그리고 `allow`를 추가하세요.

이 세 가지는 암호화된 미디어 확장을 사용해 비디오를 보호하고, 기기의 모션 및 방향 센서에 접근을 허용하며, 기기의 기본 공유 대화상자를 통해 iframe 콘텐츠를 공유할 수 있게 합니다.

# --hints--

`allow` 요소의 `iframe` 속성에는 `encrypted-media`가 값 중 하나로 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

`allow` 요소의 `iframe` 속성에는 `gyroscope`가 값 중 하나로 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

`allow` 요소의 `iframe` 속성에는 `web-share`가 값 중 하나로 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
