---
id: 68ba9c4f1688914d72876098
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

다음에 추가할 속성은 `referrerpolicy`입니다. 이 속성은 페이지가 다른 페이지와 연결될 때 얼마나 자세한 정보를 공유할지 결정하는 규칙입니다.

`referrerpolicy` 속성을 추가하고 `strict-origin-when-cross-origin`로 설정하세요. 이렇게 하면 같은 사이트에서는 전체 주소를 공유하고, 다른 사이트에서는 사이트 이름만 공유하며, 보안이 취약한 사이트에는 아무것도 공유하지 않습니다.

# --hints--

`iframe` 요소에 `referrerpolicy` 속성을 `strict-origin-when-cross-origin`로 설정해야 합니다.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
