---
id: 68ba9c4f1688914d72876093
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이제 `iframe` 요소를 만드세요. 아직 아무것도 넣지 마세요.

# --hints--

`iframe` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('iframe'))
```

`iframe` 요소는 비어 있어야 합니다.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
