---
id: 68ba9c4f1688914d72876095
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`iframe` 요소는 표시할 리소스의 URL 또는 경로를 나타내는 값을 가진 `src` 속성도 가집니다.

`src` 요소에 `https://www.youtube.com/embed/I0_951_MPE0` 값을 가진 `iframe` 속성을 추가하세요.

이 시점에서 페이지에 비디오가 표시되는 것을 볼 수 있지만, 추가해야 할 속성이 더 있습니다.

# --hints--

`iframe` 요소에 `src`로 설정된 `https://www.youtube.com/embed/I0_951_MPE0` 속성이 있어야 합니다.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
