---
id: 6735a7370e0ae93a4577c689
title: Progetta un layout per un giornale
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

In questo laboratorio, progetterai un layout per un giornale usando CSS Grid, includendo concetti come righe e colonne della griglia.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. La tua pagina dovrebbe contenere un elemento `main` con la classe `newspaper-layout`.
1. Il `.newspaper-layout` dovrebbe includere un `header` con la classe `title` contenente un `h1` per mostrare il nome del giornale.
1. Il `.newspaper-layout` dovrebbe includere un `article` con la classe `feature-article` per l'articolo principale.
1. Il `.feature-article` dovrebbe includere un elemento `h2` per il titolo dell'articolo seguito da un elemento `p` per il contenuto dell'articolo.
1. Dovresti aggiungere altri tre elementi `article` per articoli più piccoli, con classi `small-article1`, `small-article2` e `small-article3`.
1. Ciascuno degli articoli più piccoli dovrebbe includere un elemento `h3` per il titolo dell'articolo seguito da un elemento `p` per il contenuto dell'articolo.
1. Il `.newspaper-layout` dovrebbe includere un elemento `article` con la classe `secondary-article` per una sezione di notizie aggiuntiva.
1. Il `.secondary-article` dovrebbe includere un elemento `h2` per il titolo dell'articolo seguito da un elemento `p` per il contenuto dell'articolo.
1. Il `.newspaper-layout` dovrebbe includere un `figure` con la classe `cover-image` per mostrare un'immagine che rappresenta il contenuto del giornale.
1. Gli elementi con le classi `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2` e `small-article3` dovrebbero avere una proprietà `grid-area` impostata sullo stesso nome di classe.
1. Il tuo `.newspaper-layout` dovrebbe usare CSS Grid con una proprietà `grid-template-areas` per definire la disposizione delle sezioni:
   - Il `.title` dovrebbe estendersi su tutta la riga superiore.
   - Il `.feature-article` dovrebbe estendersi su due colonne nella seconda riga, con il `.cover-image` nella terza colonna.
   - Il `.secondary-article` dovrebbe estendersi su due colonne nella terza riga, con il `.cover-image` nella terza colonna.
   - I tre articoli piccoli dovrebbero apparire nella quarta riga.
1. Il `.newspaper-layout` dovrebbe avere una proprietà `grid-template-columns` che divide lo spazio in tre colonne uguali.
1. Dovresti impostare la proprietà `.newspaper-layout` del `grid-template-rows` su `auto` per la prima riga e dividere lo spazio rimanente in parti uguali per le altre righe.
1. Dovresti aggiungere uno spazio tra gli elementi della griglia.
1. Assicurati che l'immagine dentro `.cover-image` si adatti allo spazio designato impostando la sua proprietà `max-width` su `100%`.

**Nota:** Assicurati di collegare il tuo foglio di stile nel tuo HTML e applicare il CSS.

# --hints--

Dovresti avere un elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Il tuo elemento `main` dovrebbe avere una `class` di `newspaper-layout`.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

Dovresti avere un elemento `header` con la `class` di `title` all'interno del tuo elemento `.newspaper-layout`.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

Il tuo elemento `header` dovrebbe contenere un elemento `h1`.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

Il tuo elemento `.title` dovrebbe contenere del testo.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

Dovresti avere un elemento `article` per ogni articolo con classi `feature-article`, `small-article1`, `small-article2`, `small-article3` e `secondary-article`, rispettivamente.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

Il `.feature-article` dovrebbe includere un elemento `h2` seguito da un elemento `p`.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2` e `.small-article3` dovrebbero includere un elemento `h3` seguito da un elemento `p`.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

Il `.secondary-article` dovrebbe includere un elemento `h2` seguito da un elemento `p`.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

Dovresti avere un elemento `figure` con la classe `cover-image` all'interno del tuo `.newspaper-layout`.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

Il tuo `.cover-image` dovrebbe contenere un elemento `img`.

```js
assert.exists(document.querySelector(".cover-image img"));
```

La proprietà `display` dell'elemento `.newspaper-layout` dovrebbe essere impostata su `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

Il tuo elemento `.newspaper-layout` dovrebbe avere una proprietà `grid-template-areas` con `title` che si estende su tutte e tre le colonne della prima riga.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

La seconda riga del modello di griglia dovrebbe contenere un `feature-article` che si estende su due colonne e un `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

La terza riga del modello di griglia dovrebbe contenere un `secondary-article` che si estende su due colonne e un `cover-image`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

La quarta riga del modello di griglia dovrebbe contenere `small-article1`, `small-article2` e `small-article3`.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

Il tuo selettore `.newspaper-layout` dovrebbe usare la proprietà `grid-template-columns` per dividere lo spazio in tre parti uguali. Ricorda che puoi usare la funzione `repeat` oppure definire manualmente ogni unità frazionaria.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Il tuo selettore `.newspaper-layout` dovrebbe avere una proprietà `grid-template-rows` impostata su `auto 1fr 1fr 1fr`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

Dovresti aggiungere uno spazio tra gli elementi della griglia.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title` dovrebbe avere una `grid-area` di `title`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article` dovrebbe avere una `grid-area` di `feature-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image` dovrebbe avere una `grid-area` di `cover-image`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article` dovrebbe avere una `grid-area` di `secondary-article`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1` dovrebbe avere una `grid-area` di `small-article1`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2` dovrebbe avere una `grid-area` di `small-article2`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3` dovrebbe avere una `grid-area` di `small-article3`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

Il `img` dentro `.cover-image` dovrebbe avere una `max-width` di `100%`.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
