---
id: 0095d0567b6817114989a1dd
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop implementieren Sie den Algorithmus für den kürzesten Weg in JavaScript. Sie schreiben eine Funktion, die den kürzesten Weg zwischen Knoten in einem gewichteten Graphen berechnet und sowohl die Distanzen als auch die zurückgelegten Pfade zurückgibt.

Zum Beispiel findet der Algorithmus in einem Graphen, in dem Städte durch Straßen mit unterschiedlichen Entfernungen verbunden sind, die kürzeste Route von einer Stadt zur anderen. Wenn Sie von Stadt A nach Stadt D reisen möchten, könnte der Algorithmus herausfinden, dass der Weg A ⇨ B ⇨ C ⇨ D (insgesamt: 15 km) kürzer ist als der direkte Weg A ⇨ D (20 km).

Um zu beginnen, definieren Sie eine Variable namens `INF` und weisen ihr den Wert `Infinity` zu. Später verwenden Sie diese, um anzuzeigen, dass keine direkte Verbindung zwischen zwei Knoten besteht.

# --hints--

Sie sollten eine Variable namens `INF` definieren.

```js
assert.isDefined(INF);
```

Ihre `INF`-Variable sollte den Wert `Infinity` haben.

```js
assert.strictEqual(INF, Infinity);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
