---
id: 69c918841987bf16c1d1bac1
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Jetzt fügen Sie die Hauptschleife hinzu, die den Algorithmus steuert. Sie läuft einmal für jeden Knoten im Graphen und wählt jedes Mal den nächstgelegenen unbesuchten Knoten aus.

Fügen Sie innerhalb der `shortestPath`-Funktion eine `for`-Schleife hinzu, die `n`-mal mit `i` als Schleifenvariable läuft.

Am Anfang des Schleifenkörpers deklarieren Sie mit `let` zwei Variablen `minDistance` und `current`, die auf `INF` beziehungsweise `-1` gesetzt sind. Diese verfolgen den nächstgelegenen unbesuchten Knoten, der in jeder Iteration gefunden wird.

# --hints--

Sie sollten eine `for`-Schleife innerhalb von `shortestPath` haben, die `n`-mal läuft. Vergessen Sie nicht, `i` mit `let` darin zu setzen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+i\s*=\s*0\s*;\s*i\s*<\s*n\s*;\s*i\+\+\s*\)/
);
```

Sie sollten eine `minDistance`-Variable haben, die innerhalb der Schleife auf `INF` gesetzt ist.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+minDistance\s*=\s*INF/
);
```

Sie sollten eine `current`-Variable haben, die innerhalb der Schleife auf `-1` gesetzt ist.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(let|var)\s+current\s*=\s*-1/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
