---
id: 69c937c73f23b1140bcd297a
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Der Parameter `targetNode` ist optional. Wenn er nicht angegeben wird, berechnet die Funktion die kürzesten Pfade von `startNode` zu allen anderen Knoten im Graphen.

Geben Sie dem Parameter `targetNode` einen Standardwert von `null`.

# --hints--

Der Parameter `targetNode` Ihrer `shortestPath`-Funktion sollte standardmäßig auf `null` gesetzt sein.

```js
assert.match(
  __helpers.removeJSComments(code),
  /targetNode\s*=\s*null/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode) {
--fcc-editable-region--
  
}
```
