---
id: 69c918841987bf16c1d1baba
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Sie benötigen ein 2D-Array, um die Adjazenzmatrix des Graphen darzustellen. Jeder Wert repräsentiert das Gewicht (die Entfernung) der Kante zwischen zwei Knoten. Ein Wert von `INF` bedeutet, dass keine direkte Kante existiert.

Erstellen Sie eine weitere Variable mit dem Namen `adjMatrix` und weisen Sie ihr ein 2D-Array mit den folgenden Werten zu:

```js
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Sie sollten eine Variable mit dem Namen `adjMatrix` erstellen.

```js
assert.isDefined(adjMatrix);
```

`adjMatrix` sollte ein 2D-Array mit 6 Zeilen und 6 Spalten sein.

```js
assert.isArray(adjMatrix);
assert.lengthOf(adjMatrix, 6);
adjMatrix.forEach(row => assert.lengthOf(row, 6));
```

`adjMatrix` sollte die korrekten Werte enthalten.

```js
assert.deepEqual(adjMatrix[0], [0, 5, 3, INF, 11, INF]);
assert.deepEqual(adjMatrix[1], [5, 0, 1, INF, INF, 2]);
assert.deepEqual(adjMatrix[2], [3, 1, 0, 1, 5, INF]);
assert.deepEqual(adjMatrix[3], [INF, INF, 1, 0, 9, 3]);
assert.deepEqual(adjMatrix[4], [11, INF, 5, 9, 0, INF]);
assert.deepEqual(adjMatrix[5], [INF, 2, INF, 3, INF, 0]);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
--fcc-editable-region--

--fcc-editable-region--
```
