---
id: 69c918841987bf16c1d1babd
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Sie müssen die kürzeste bekannte Entfernung vom Startknoten zu jedem anderen Knoten verfolgen. Zu Beginn nehmen Sie an, dass jeder Knoten unendlich weit entfernt ist.

In JavaScript können Sie ein Array erstellen, das mit einem Wert vorab gefüllt ist, indem Sie `new Array(n).fill(value)` verwenden. Zum Beispiel erstellt `new Array(3).fill(0)` `[0, 0, 0]`.

Erstellen Sie innerhalb der `shortestPath`-Funktion eine Variable namens `distances` und weisen Sie ihr `new Array(n).fill(INF)` zu.

# --hints--

Sie sollten innerhalb von `distances` eine Variable namens `shortestPath` erstellen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distances\s*=/
);
```

Ihre `distances`-Variable sollte auf `new Array(n).fill(INF)` gesetzt werden.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*INF\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--
function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;

}
--fcc-editable-region--
```
