---
id: 69c918841987bf16c1d1bacb
title: Schritt 20
challengeType: 1
dashedName: step-20
---

# --description--

Sobald die Hauptschleife beendet ist, müssen Sie entscheiden, für welchen Knoten bzw. welche Knoten die Ergebnisse angezeigt werden sollen.

Wenn ein bestimmter `targetNode` angegeben wurde, zeigen Sie nur die Ergebnisse für diesen Knoten an. Andernfalls zeigen Sie die Ergebnisse für alle Knoten an.

In JavaScript können Sie alle Indizes eines Arrays als iterierbares Objekt mit `[...Array(n).keys()]` erhalten.

Nach der äußeren `for`-Schleife erstellen Sie eine Variable namens `targets`. Verwenden Sie einen ternären Ausdruck: Wenn `targetNode` strikt nicht `null` ist, weisen Sie `[targetNode]` zu, andernfalls `[...Array(n).keys()]`.

# --hints--

Sie sollten nach der Hauptschleife eine Variable `targets` erstellen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+targets\s*=/
);
```

Ihre `targets`-Variable sollte einen ternären Ausdruck verwenden, der `targetNode !== null` überprüft.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /targets\s*=\s*targetNode\s*!==\s*null\s*\?/
);
```

Wenn ein Ziel angegeben ist, sollte `targets` `[targetNode]` sein, andernfalls `[...Array(n).keys()]`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /\?\s*\[\s*targetNode\s*\]\s*:\s*\[\s*\.\.\.\s*Array\s*\(\s*n\s*\)\s*\.keys\s*\(\s*\)\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
--fcc-editable-region--

--fcc-editable-region--
}
```
