---
id: 69c918841987bf16c1d1bac7
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Da der aktuelle Knoten nun als besucht markiert ist, müssen Sie alle seine Nachbarn betrachten, um zu prüfen, ob Sie sie effizienter über den aktuellen Knoten erreichen können.

Fügen Sie nach `visited[current] = true` eine Nachbar-`for`-Schleife hinzu, die durch jeden `nodeNo` von `0` bis `n` iteriert. Beachten Sie, dass dies eine weitere innere Schleife auf derselben Ebene wie die erste ist.

Deklarieren Sie innerhalb der Schleife eine Variable namens `distance` und weisen Sie ihr `matrix[current][nodeNo]` zu. So erhalten Sie das Kantengewicht zwischen dem aktuellen Knoten und seinem Nachbarn.

# --hints--

Sie sollten nach dem Markieren des aktuellen Knotens als besucht eine Nachbar-`for`-Schleife hinzufügen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(/
);
```

Die Nachbar-`for`-Schleife sollte von `0` bis `n` iterieren.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /visited\s*\[\s*current\s*\]\s*=\s*true[\s\S]*?for\s*\(\s*(var|let)\s+\w+\s*=\s*0\s*;\s*\w+\s*<\s*n\s*;\s*\w+\+\+\s*\)/
);
```

Innerhalb der Nachbarschaftsschleife sollten Sie eine `distance`-Variable haben, der `matrix[current][nodeNo]` zugewiesen ist.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+distance\s*=\s*matrix\s*\[\s*current\s*\]\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
