---
id: 69c918841987bf16c1d1babf
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Sie müssen auch den tatsächlichen Pfad verfolgen, der zu jedem Knoten führt. Sie speichern eine Liste von Pfaden, wobei jeder Eintrag ein Array von Knotenindizes ist, das die zurückgelegte Route darstellt.

Anfangs enthält der Pfad jedes Knotens nur sich selbst. Sie können diese Struktur mit `Array.from()` erstellen:

```js
Array.from({ length: n }, (_, i) => [i]);
```

Dies erzeugt ein Array mit `n` Elementen, wobei das Element am Index `i` `[i]` ist.

Erstellen Sie noch innerhalb der `shortestPath`-Funktion eine Variable namens `paths` und weisen Sie ihr `Array.from({ length: n }, (_, i) => [i])` zu.

# --hints--

Sie sollten eine Variable `paths` innerhalb der `shortestPath`-Funktion erstellen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+paths\s*=/
);
```

Ihre `paths`-Variable sollte mit `Array.from` initialisiert werden.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*=\s*Array\.from/
);
```

Ihre `paths`-Variable sollte `{ length: n }` als erstes Argument an `Array.from` übergeben.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\(\s*\{\s*length\s*:\s*n\s*\}/
);
```

Das zweite Argument von `Array.from` sollte ein Callback `(_, i) => [i]` sein, wobei `_` ein Platzhalter für das ungenutzte Element und `i` der Index ist. Es sollte `[i]` zurückgeben.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /Array\.from\s*\([\s\S]*?,\s*function\s*\(\s*\w+\s*,\s*(\w+)\s*\)\s*\{[\s\S]*?return\s*\[\s*\1\s*\]|Array\.from\s*\([\s\S]*?,\s*\(\s*\w+\s*,\s*\w+\s*\)\s*=>\s*\[\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
