---
id: 69c918841987bf16c1d1bac5
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Nachdem Sie alle Knoten durchsucht haben, um den nächstgelegenen unbesuchten zu finden, müssen Sie den Fall behandeln, dass kein gültiger Knoten gefunden wird. Dies passiert, wenn alle verbleibenden Knoten unerreichbar sind.

Fügen Sie nach der inneren `for`-Schleife, aber noch innerhalb der äußeren Schleife, eine `if`-Anweisung hinzu, die überprüft, ob `current` strikt gleich `-1` ist. Wenn das zutrifft, verwenden Sie `break` innerhalb des `if`-Blocks, um die äußere Schleife vorzeitig zu beenden.

# --hints--

Sie sollten eine `if`-Anweisung hinzufügen, die überprüft, ob `current === -1`.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)/
);
```

Sie sollten `break` innerhalb des `if`-Blocks verwenden.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*current\s*===\s*-1\s*\)[\s\S]*?break/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
