---
id: 69c918841987bf16c1d1bac2
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Innerhalb der Schleife müssen Sie jeden Knoten durchsuchen, um denjenigen mit der kleinsten bekannten Entfernung zu finden, der noch nicht besucht wurde.

Fügen Sie eine `for`-Schleife innerhalb der Hauptschleife hinzu, die durch jeden Knoten iteriert. Verwenden Sie `nodeNo` als Schleifenvariable, die von `0` bis `n` läuft.

Lassen Sie den Rumpf vorerst leer.

# --hints--

Sie sollten eine zweite (innere) `for`-Schleife innerhalb der Hauptschleife haben, mit `nodeNo` als Schleifenvariable.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /for\s*\(\s*(let|var)\s+nodeNo\s*=\s*0\s*;\s*nodeNo\s*<\s*n\s*;\s*nodeNo\+\+\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
