---
id: 69c918841987bf16c1d1bac0
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Während der Algorithmus läuft, müssen Sie verfolgen, welche Knoten bereits verarbeitet wurden, damit Sie sie nicht erneut besuchen.

Erstellen Sie eine Variable namens `visited` und weisen Sie ihr `new Array(n).fill(false)` zu.

# --hints--

Sie sollten eine Variable namens `visited` erstellen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+visited\s*=/
);
```

Ihre `visited`-Variable sollte auf `new Array(n).fill(false)` gesetzt werden.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /new\s+Array\s*\(\s*n\s*\)\s*\.fill\s*\(\s*false\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
--fcc-editable-region--

--fcc-editable-region--
}
```
