---
id: 69c918841987bf16c1d1bac8
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Bevor Sie die Distanzen aktualisieren, müssen Sie überprüfen, ob der Nachbar überhaupt berücksichtigt werden sollte. Es muss eine tatsächliche Kante zu ihm geben (`distance !== INF`) und er darf noch nicht besucht worden sein.

Fügen Sie innerhalb der Nachbar-`for`-Schleife eine `if`-Anweisung hinzu, die überprüft, ob `distance` nicht strikt gleich `INF` **und** ob `nodeNo` nicht besucht ist.

Deklarieren Sie innerhalb des `if`-Blocks eine Variable `newDistance`, der der Wert `distances[current] + distance` zugewiesen wird. Dies sind die Gesamtkosten, um den Nachbarn über den aktuellen Knoten zu erreichen.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `distance !== INF && !visited[nodeNo]` hinzufügen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*distance\s*!==\s*INF\s*&&\s*!\s*visited\s*\[\s*\w+\s*\]\s*\)/
);
```

Sie sollten innerhalb des `newDistance`-Blocks eine `distances[current] + distance` mit dem Wert `if` deklarieren.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+newDistance\s*=\s*distances\s*\[\s*current\s*\]\s*\+\s*distance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      --fcc-editable-region--

      --fcc-editable-region--
    }
  }
}
```
