---
id: 69c918841987bf16c1d1bace
title: Schritt 23
challengeType: 1
dashedName: step-23
---

# --description--

Da Sie nun wissen, dass der Knoten erreichbar ist und nicht der Start, formatieren Sie seinen Pfad in einen lesbaren String.

Das `paths[nodeNo]`-Array enthält die Folge der besuchten Knotenindizes, um `nodeNo` zu erreichen. Sie können daraus mit der `0 -> 2 -> 3`-Methode einen für Menschen lesbaren String wie `.join()` erstellen.

Innerhalb der `for...of`-Schleife, nach der `if`-Anweisung mit `continue`, deklarieren Sie eine Variable namens `path` und weisen ihr den Wert `paths[nodeNo].join(' -> ')` zu.

# --hints--

Sie sollten eine Variable namens `path` innerhalb der `for...of`-Schleife haben.

```js
assert.match(
  __helpers.removeJSComments(code),
  /for\s*\(\s*(const|let)\s+nodeNo\s+of\s+targets\s*\)[\s\S]*?(const|let)\s+path\s*=/
);
```

Ihre `path`-Variable sollte den Wert `paths[nodeNo].join(' -> ')` haben.

```js
assert.match(
  __helpers.removeJSComments(code),
  /path\s*=\s*paths\s*\[\s*\w+\s*\]\s*\.\s*join\s*\(\s*['"]\s*->\s*['"]\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);
  
  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;
    
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          paths[nodeNo] = [...paths[current], nodeNo];
        }
      }
    }
  }
  
  const targets = targetNode !== null ? [targetNode] : [...Array(n).keys()];
  for (const nodeNo of targets) {
    if (nodeNo === startNode || distances[nodeNo] === INF) {
      continue;
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
  }
}
```
