---
id: 69c918841987bf16c1d1baca
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Wenn Sie einen kürzeren Weg zu einem Nachbarn finden, müssen Sie auch den gespeicherten Pfad aktualisieren, um ihn zu erreichen.

In JavaScript können Sie mit dem Spread-Operator ein neues Array erstellen, das ein bestehendes Array und ein neues Element kombiniert:

```js
const newPath = [...existingArray, newElement];
```

Innerhalb des verschachtelten `if`-Blocks aktualisieren Sie `paths[nodeNo]` so, dass es der Pfad zum aktuellen Knoten mit `nodeNo` am Ende ist, indem Sie die Spread-Syntax wie folgt verwenden: `[...paths[current], nodeNo]`.

# --hints--

Sie sollten `paths[nodeNo]` innerhalb des verschachtelten `if`-Blocks aktualisieren.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=/
);
```

`paths[nodeNo]` sollte auf `[...paths[current], nodeNo]` gesetzt werden.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /paths\s*\[\s*\w+\s*\]\s*=\s*\[\s*\.\.\.\s*paths\s*\[\s*current\s*\]\s*,\s*\w+\s*\]/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        if (newDistance < distances[nodeNo]) {
          distances[nodeNo] = newDistance;
          --fcc-editable-region--

          --fcc-editable-region--
        }
      }
    }
  }
}
```
