---
id: 69c918841987bf16c1d1bac9
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

Jetzt sollten Sie prüfen, ob der Weg über den aktuellen Knoten eine kürzere Route zum Nachbarn ergibt. Wenn `newDistance` besser ist als das, was bereits gespeichert ist, sollten Sie es aktualisieren.

Fügen Sie innerhalb des bestehenden `if`-Blocks eine verschachtelte `if`-Anweisung hinzu, die überprüft, ob `newDistance` kleiner ist als `distances[nodeNo]`. Aktualisieren Sie innerhalb dieses verschachtelten `if`-Blocks `distances[nodeNo]` auf `newDistance`.

# --hints--

Sie sollten eine verschachtelte `if`-Anweisung mit der Bedingung `newDistance < distances[nodeNo]` hinzufügen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /if\s*\(\s*newDistance\s*<\s*distances\s*\[\s*\w+\s*\]\s*\)/
);
```

Sie sollten `distances[nodeNo]` innerhalb des verschachtelten `newDistance`-Blocks auf `if` aktualisieren.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /distances\s*\[\s*\w+\s*\]\s*=\s*newDistance/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
  const n = matrix.length;
  const distances = new Array(n).fill(INF);
  distances[startNode] = 0;
  const paths = Array.from({ length: n }, (_, i) => [i]);
  const visited = new Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    let minDistance = INF;
    let current = -1;
    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      if (!visited[nodeNo] && distances[nodeNo] < minDistance) {
        minDistance = distances[nodeNo];
        current = nodeNo;
      }
    }
    if (current === -1) {
      break;
    }
    visited[current] = true;

    for (let nodeNo = 0; nodeNo < n; nodeNo++) {
      const distance = matrix[current][nodeNo];
      if (distance !== INF && !visited[nodeNo]) {
        const newDistance = distances[current] + distance;
        --fcc-editable-region--

        --fcc-editable-region--
      }
    }
  }
}
```
