---
id: 69c918841987bf16c1d1babc
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Sie müssen die Anzahl der Knoten im Graphen speichern.

Erstellen Sie innerhalb der `shortestPath`-Funktion eine Variable namens `n` und weisen Sie ihr `matrix.length` zu.

# --hints--

Sie sollten innerhalb der `shortestPath`-Funktion eine Variable `n` erstellen.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=/
);
```

Ihre `n`-Variable sollte den Wert `matrix.length` haben.

```js
assert.match(
  __helpers.removeJSComments(shortestPath.toString()),
  /(const|let|var)\s+n\s*=\s*matrix\.length/
);
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

function shortestPath(matrix, startNode, targetNode = null) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
