---
id: 69c918841987bf16c1d1babb
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Jetzt erstellen Sie die Hauptfunktion, die die gesamte Arbeit erledigt.

Erstellen Sie eine Funktion namens `shortestPath`, die drei Parameter entgegennimmt: `matrix`, `startNode` und `targetNode`.

In den nächsten Schritten bauen Sie die Logik für diese Funktion auf.

# --hints--

Sie sollten eine Funktion namens `shortestPath` erstellen.

```js
assert.isFunction(shortestPath);
```

Ihre `shortestPath`-Funktion sollte die Parameter `matrix`, `startNode` und `targetNode` haben.

```js
const params = __helpers.getFunctionParams(shortestPath.toString());
const names = params.map(p => p.name);
assert.include(names, 'matrix');
assert.include(names, 'startNode');
assert.include(names, 'targetNode');
```

# --seed--

## --seed-contents--

```js
const INF = Infinity;
const adjMatrix = [
  [0, 5, 3, INF, 11, INF],
  [5, 0, 1, INF, INF, 2],
  [3, 1, 0, 1, 5, INF],
  [INF, INF, 1, 0, 9, 3],
  [11, INF, 5, 9, 0, INF],
  [INF, 2, INF, 3, INF, 0],
];

--fcc-editable-region--

--fcc-editable-region--
```
