---
id: 69b83e35f19ba26ba1fa517a
title: 신호등 순서 제어기 만들기
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

이 실습에서는 구성 가능한 신호등 주기를 시뮬레이션하고 이상 현상을 기록할 거예요.

신호등의 단계를 설명하는 구성 객체를 다룰 거예요. 각 구성 객체는 다음 속성을 가집니다:

- `fault`: `true`일 때 조기 종료를 유발하는 불리언 플래그입니다.
- `phases`: 단계 객체들의 배열입니다.

`phases` 안의 각 단계 객체는 다음 속성을 가집니다:

- `color`: 신호등 색상을 나타내는 문자열(`"green"`, `"yellow"`, 또는 `"red"`).
- `duration`: 단계가 지속되는 시간을 초 단위로 나타내는 양의 정수입니다.

제공된 `config1`, `config2`, `config3`, `config4` 객체를 가능한 구성 객체 예시로 참고하세요.

**참고**: 테스트 실패를 방지하려면 추가 `console.log()` 문을 넣지 마세요.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

# 사용자 스토리

1. `runSequence`라는 이름의 함수를 만들고 두 개의 매개변수 `config`와 `cycles`를 받으세요. 여기서 `config`는 구성 객체를, `cycles`는 시퀀스가 멈추기 전 최대 실행 횟수를 나타냅니다.

1. `runSequence(config, cycles)`를 `for` 또는 `while` 루프를 사용해 주어진 횟수만큼 각 단계를 순차적으로 반복하도록 구현하세요.

1. `runSequence` 함수는 다음을 수행해야 합니다:
    - `config.phases`가 비어 있으면 `No phases found`를 기록하고 즉시 반환하세요.
    - `config.fault`가 `true`이면 `Faulted phase!`를 기록하고 시뮬레이션을 조기 종료하세요.
    - `duration <= 0`이면 `Invalid phase detected`를 기록하세요.
    - 각 유효한 단계에 대해 `Switching to [color] for [duration] s`를 기록하세요. 여기서 `[color]`와 `[duration]`은 단계 객체의 해당 속성으로 대체합니다.
    - 예를 들어, `runSequence(config1, 1)`은 다음을 기록해야 합니다:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. `generateTimeline`이라는 이름의 함수를 만들고 두 개의 매개변수 `config`와 `cycles`를 받으세요.

1. `generateTimeline` 함수는 다음을 수행해야 합니다:
    - 각 단계가 끝난 누적 경과 시간을 배열에 기록하세요. 반복하면서 각 단계의 `duration`을 누적 합계에 더하세요.
    - `config.fault`가 `true`이거나 `duration <= 0`이어도 모든 결함 및 유효하지 않은 단계를 검증 없이 처리하세요.
    - 누적 타임스탬프 배열을 반환하세요.
    - 예를 들어, `generateTimeline(config1, 1)`은 `[5, 7, 11]` 배열을 반환해야 합니다.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

`config1`, `config2`, `config3`, `config4` 객체는 변경하지 마세요.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

`runSequence`라는 이름의 함수를 만들어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.runSequence);
```

`runSequence` 함수는 두 개의 매개변수 `config`와 `cycles`를 받아야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.runSequence?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`runSequence(config1, 1)`은 순서대로 `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`를 기록해야 합니다.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)`은 순서대로 `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`를 기록해야 합니다.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)`은 순서대로 `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s`를 기록해야 합니다.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)`는 `Faulted phase!`를 기록해야 합니다.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)`는 `No phases found`를 기록해야 합니다.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

`generateTimeline`이라는 이름의 함수를 만들어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.generateTimeline);
```

`generateTimeline` 함수는 두 개의 매개변수 `config`와 `cycles`를 받아야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.generateTimeline?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`generateTimeline(config1, 1)`은 `[5, 7, 11]`을 반환해야 합니다.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)`은 `[5, 7, 11, 16, 18, 22]`를 반환해야 합니다.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)`은 `[3, 1, 7, 10, 8, 14]`를 반환해야 합니다.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)`은 `[5, 7, 13]`을 반환해야 합니다.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)`은 `[]`를 반환해야 합니다.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
