---
id: 6894c1b715b8d18a18ef3203
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、Pythonのクラスとオブジェクトを使って楽器の在庫管理を作成しながら練習します。

Pythonでクラスを定義するには、`class` キーワードの後にクラス名とコロンを続けて書きます。

```python
class ClassName:
    # class contents go here
```

Pythonでは、クラス名は通常PascalCaseを使います。これは単語の最初の文字が大文字で、単語の間にスペースを入れない書き方です。

まず、`MusicalInstrument` というクラスを作成してください。今はクラスの本体として`pass` を追加しましょう。

# --hints--

`MusicalInstrument` というクラスを作成してください。

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
