---
id: 6895e3f91c069e0068a42582
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

`play` メソッドの中で、`print()` を使って楽器を演奏することについてのメッセージを表示してください。

f文字列を使って、メッセージに楽器の名前を含めてください：`The [instrument name] is fun to play!`。`[instrument name]` は実際の楽器の名前に置き換えられます。

# --hints--

`play` メソッドでは、f文字列を使ってメッセージ `The [instrument name] is fun to play!` を表示してください。楽器の名前にアクセスするには `self.name` を使います。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
