---
id: 6895dffc64e97df0e62e9ef9
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

メソッドはクラス内で定義される関数で、クラスのデータを使って動作を実行できます。メソッドの作り方は次の通りです。

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

Pythonでは、`__init__` メソッドはイニシャライザーメソッドと呼ばれる特別なメソッドで、クラスの新しいインスタンスを作成するときに呼び出されます。

`MusicalInstrument` クラスの中に、`__init__` メソッドを3つのパラメータ `self`、`name`、`instrument_type` で追加してください。

`self` パラメータはクラスのインスタンスを指す標準的な名前で、すべてのインスタンスメソッドの最初のパラメータとして必要です。

# --hints--

`__init__` クラスの中に`MusicalInstrument` メソッドを定義してください。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

`__init__` メソッドには3つのパラメータ `self`、`name`、`instrument_type` を持たせてください。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    pass
--fcc-editable-region--
```
