---
id: 6895e0452b3370f26f365a46
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

`__init__` メソッドでは、パラメータをインスタンス属性に代入する必要があります。インスタンス属性は、クラスの特定のインスタンスに属する変数です。

`__init__` メソッド内で、`name` パラメータを `self.name` に、`instrument_type` パラメータを `self.instrument_type` に代入してください。

# --hints--

`name` メソッド内で `self.name` パラメータを `__init__` に代入する必要があります。

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

`instrument_type` メソッド内で `self.instrument_type` パラメータを `__init__` に代入する必要があります。

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        pass
--fcc-editable-region--
```
