---
id: 6895e4940e39950288e406c2
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

次に、`get_fact` という名前のメソッドをクラスに追加して、その楽器に関する事実を返してください。

このメソッドは `self` パラメータのみを持ち、`The [instrument name] is part of the [instrument type] family of instruments.` と言う f-string を返す必要があります。

メッセージを直接出力する `play` メソッドとは異なり、このメソッドはコードの他の場所で使える文字列を返すべきです。

# --hints--

`get_fact` という名前のメソッドを `self` パラメータで定義してください。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

`get_fact` メソッドは楽器の名前と種類を含む f-string を返すべきです。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
