---
id: 6895e235935106f88281b648
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

インスタンスの属性にはドット記法でアクセスできます。

```py
instance_name.attribute_name
```

これにより、初期化時にストアされた値を取得できます。

両方のインスタンスの`name`と`instrument_type`属性を出力してみてください。作成時に代入した値が確認できます。

# --hints--

`name`の`instrument_1`属性を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

`instrument_type`の`instrument_1`属性を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

`name`の`instrument_2`属性を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

`instrument_type`の`instrument_2`属性を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
```
