---
id: 6895e356fd877efd6e3657ee
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--


`MusicalInstrument` クラスの中に、`play` パラメータだけを持つ `self` という名前の別のメソッドを定義してください。`self` パラメータは、メソッドが特定のインスタンスの属性にアクセスできるようにします。

# --hints--

`play` パラメータを持つ `self` という名前のメソッドを定義してください。

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

```
