---
id: 6895e0bdb0348bf46952ab67
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

いよいよ `MusicalInstrument` クラスのインスタンスを作成する時です。

クラスのインスタンスを作成するには、クラス名を関数のように呼び出し、`__init__` メソッドに必要な引数を渡します（`self` は Python が自動的に処理します）。

```python
my_instance = ClassName(arg1, arg2)
```

このインスタンスを作成すると、Python は `__init__` メソッドを呼び出し、その値をこの特定のインスタンスだけに属するインスタンス属性として保存します。

クラス定義の外側で、`instrument_1` 楽器である `Oboe` の `woodwind` という名前のインスタンスを作成してください。

# --hints--

`MusicalInstrument` クラスの `instrument_1` という名前のインスタンスを作成してください。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

`Oboe` を作成するときに `woodwind` と `instrument_1` を引数として渡してください。

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
```
