---
id: 6895e197180c9af6988e1e19
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

`instrument_2` という名前の2つ目のインスタンスを、`Trumpet` である `brass` 楽器のために作成してください。

同じクラスのテンプレートから複数のインスタンスを作成できることに注目してください。各インスタンスは、他のインスタンスとは独立して自分の属性の値を保持します。

# --hints--

`instrument_2` という名前の `MusicalInstrument` クラスのインスタンスを作成してください。

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

`Trumpet` を作成するときに `brass` と `instrument_2` を引数として渡してください。

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')

--fcc-editable-region--
```
