---
id: 6895e2becb7890fad818ad4d
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

オブジェクトの属性にアクセスする方法がわかったので、既存の `print` の呼び出しを削除してください。

# --hints--

`print(instrument_1.name)` の文を削除してください。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

`print(instrument_1.instrument_type)` の文を削除してください。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

`print(instrument_2.name)` の文を削除してください。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

`print(instrument_2.instrument_type)` の文を削除してください。

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--
print(instrument_1.name)
print(instrument_1.instrument_type)

print(instrument_2.name)
print(instrument_2.instrument_type)

--fcc-editable-region--
```
