---
id: 6895e5337b1914050b0af7bd
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

今、`get_fact()` インスタンスの `instrument_1` メソッドを呼び出して、その結果を出力してください。

`get_fact()` は直接出力するのではなく文字列を返すので、返された文字列を表示するには `print()` を使う必要があります。

# --hints--

`get_fact()` の `instrument_1` を呼び出して、その結果を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()

--fcc-editable-region--
```
