---
id: 6895e572df8596062234f571
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

最後に、`instrument_2` に対しても同じことをしてください。`play()` の `instrument_2` メソッドを呼び出し、その後 `get_fact()` の `instrument_2` を呼び出した結果を出力してください。

これであなたの楽器ワークショップは完成です！

# --hints--

`play()` インスタンスの `instrument_2` メソッドを呼び出す必要があります。

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_2.play()")`))
});
```

`get_fact()` に対して `instrument_2` を呼び出し、その結果を出力してください。

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


--fcc-editable-region--
instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()
print(instrument_1.get_fact())

--fcc-editable-region--
```

# --solutions--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'


instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')


instrument_1.play()              # The Oboe is fun to play!
print(instrument_1.get_fact())   # The Oboe is part of the woodwind family of instruments.

instrument_2.play()              # The Trumpet is fun to play!
print(instrument_2.get_fact())   # The Trumpet is part of the brass family of instruments.
```
