---
id: 687ad7be1ff45032ccf1d92f
title: Construye un Theme Switcher
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

En este laboratorio, construirás una app que cambia entre diferentes temas. Cuando cambies a un tema diferente, el color de fondo cambiará y se mostrará un mensaje distinto.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un elemento `button` con el texto `Switch Theme`.
2. Your `button` element should have the following attributes:
   - Un atributo `id` establecido en `"theme-switcher-button"`.
   - Un atributo `aria-haspopup` establecido en `"true"`.
   - Un atributo `aria-expanded` establecido en `"false"`.
   - Un atributo `aria-controls` establecido en `"theme-dropdown"`.
3. You should have a `ul` element with the following attributes:
   - Un atributo `id` establecido en `"theme-dropdown"`.
   - Un atributo `role` configurado como `"menu"`.
   - Un atributo `aria-labelledby` establecido en `"theme-switcher-button"`.
   - Un atributo `hidden`.
4. Tu elemento `ul` debe tener al menos dos elementos `li` con un atributo `role` establecido en `"menuitem"` y un texto de tu elección que represente un tema diferente.
5. Cada uno de tus elementos `li` debe tener un atributo `id` que comience con `theme-` y termine con el tema que establezcas para el texto del elemento `li`. Por ejemplo, si uno de tus temas es `Light`, entonces tu `id` debe ser `theme-light`.
6. Debes tener un elemento con un atributo `aria-live` establecido en `"polite"`.
7. Debes tener un arreglo `themes` con al menos dos objetos que contengan cada uno una propiedad `name` y `message`. El `name` representará un tema diferente y el `message` se mostrará cuando el tema cambie. Puedes elegir los valores de `name` y `message` que prefieras, pero los valores de `name` deben coincidir con una de las opciones que configures en tus elementos `li`.
8. Cuando un usuario hace clic en el `#theme-switcher-button`, el elemento `#theme-dropdown` debe mostrar los temas disponibles.
9. Cuando se muestra el elemento `#theme-dropdown`, el atributo `aria-expanded` debe establecerse en `"true"` para el elemento `button`.
10. Cuando el usuario hace clic en el `#theme-switcher-button` mientras el elemento `#theme-dropdown` está visible, el menú debe cerrarse, debe añadirse el atributo `hidden` y el atributo `aria-expanded` debe establecerse en `"false"`.
11. Cuando un usuario hace clic en el `#theme-switcher-button` y selecciona un tema, la clase correspondiente `theme-<name>` debe añadirse al elemento `body` y el mensaje del tema relacionado que configures en el arreglo `themes` debe mostrarse en el elemento con `aria-live="polite"`.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

Debería tener un elemento `button`.

```js
assert.exists(document.querySelector("button"));
```

Tu elemento `button` debe tener el texto `Switch Theme`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

Tu elemento `button` debe tener un atributo `id` establecido en `"theme-switcher-button"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

Tu elemento `button` debe tener un atributo `aria-haspopup` establecido en `"true"`. 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

Tu elemento `button` debe tener un atributo `aria-expanded` establecido en `"false"`. 

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

Tu elemento `button` debe tener un atributo `aria-controls` establecido en `"theme-dropdown"`.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

Debes tener un elemento `ul`.

```js
assert.exists(document.querySelector("ul"));
```

Tu elemento `ul` debe tener un atributo `id` establecido en `"theme-dropdown"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

Tu elemento `ul` debe tener un atributo `role` establecido en `"menu"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

Tu elemento `ul` debe tener un atributo `aria-labelledby` establecido en `"theme-switcher-button"`.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

Tu elemento `ul` debe tener un atributo `hidden`.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

Tu elemento `ul` debe tener al menos dos elementos `li`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

Cada uno de tus elementos `li` debe tener un atributo `role` establecido en `"menuitem"`.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

Cada uno de tus elementos `li` debe tener un texto que represente un tema.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

Cada uno de tus elementos `li` debe tener un atributo `id` que comience con `theme-` y termine con el tema correspondiente para ese elemento. Asegúrate de que el valor de tu `id` esté todo en minúsculas.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

Debes tener un elemento con un atributo `aria-live` establecido en `"polite"`.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

Debes tener un arreglo `themes`.

```js
assert.isArray(themes);
```

Tu arreglo `themes` debe tener al menos dos objetos.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

Tu arreglo `themes` debe tener al menos dos objetos, cada uno con una propiedad `name` y `message`. 

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

Los valores de las propiedades `name` y `message` deben ser cadenas que no estén vacías.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

Cada uno de los valores de tu propiedad `name` debe coincidir con uno de los valores que proporcionaste para los elementos `li`. Asegúrate de que todos estos valores estén en minúsculas.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

Cuando un usuario hace clic en el `#theme-switcher-button` para mostrar las opciones de tema, se debe eliminar el atributo `hidden` del elemento `#theme-dropdown`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

Cuando se muestra el elemento `#theme-dropdown`, el atributo `aria-expanded` debe establecerse en `"true"` para el elemento `button`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

Cuando el usuario hace clic en el `#theme-switcher-button` mientras el elemento `#theme-dropdown` está visible, el menú desplegable debe volver a estar en `hidden`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

Cuando el usuario hace clic en el `#theme-switcher-button` mientras el elemento `#theme-dropdown` está visible, el atributo `aria-expanded` debe establecerse en `"false"`.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

Cuando un usuario hace clic en el `#theme-switcher-button` y selecciona un tema, la clase correspondiente `theme-<name>` debe añadirse al elemento `body`.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

Cuando un usuario hace clic en el `#theme-switcher-button` y selecciona un tema, se debe mostrar un `message` relacionado con el `theme` seleccionado del arreglo `themes` en el elemento con `aria-live="polite"`.


```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
