---
id: 61fd6b7c83dbf54a08cf0498
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Усередині вашого елемента `div` додайте три елементи `span`. Кожному з них надайте атрибут `class` зі значенням `year`, а також додайте наступний текст (у порядку): `2019`, `2020` і `2021`.

# --hints--

Ваш елемент `div` має містити три елементи `span`.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

Кожен елемент `span` повинен мати атрибут `class` зі значенням `year`.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

Перший ваш елемент `span` має містити текст `2019`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

Другий ваш елемент `span` має містити текст `2020`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

Третій ваш елемент `span` має містити текст `2021`.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
