---
id: 61fd719788899952e67692b9
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Елементи `thead` і `tbody` використовуються, щоб позначити, яка частина вашої таблиці є заголовком, а яка містить основні дані або вміст.

Додайте `thead` і `tbody` до вашого першого `table`, під елементом `caption`.

# --hints--

Ваш перший елемент `table` має містити елемент `thead`.

```js
const table = document.querySelectorAll('table')?.[0];
assert.isNotNull(table?.querySelector('thead'));
```

Ваш перший елемент `table` має містити елемент `tbody`.

```js
const table = document.querySelectorAll('table')?.[0];
const tbody = table?.querySelector('tbody');
assert.isNotNull(tbody);
```

Ваш елемент `thead` має бути безпосередньо під вашим елементом `caption`.

```js
assert.strictEqual(
  document.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

Ваш елемент `tbody` має бути безпосередньо під вашим елементом `thead`.

```js
assert.strictEqual(
  document.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
