---
id: 61fd6cc9475a784b7776233e
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Під вашим існуючим елементом `div` додайте новий елемент `div` з множиною `class`, встановленою на `table-wrap`. Це буде контейнер для ваших таблиць.

Всередині нього додайте три елементи `table`. Ви будете використовувати CSS, щоб стилізувати їх як одну таблицю, але використання окремих таблиць допоможе читачам екрана краще розуміти потік документа.

# --hints--

Ви повинні створити новий елемент `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Ваш новий елемент `div` має мати множину `class`, встановлену на `table-wrap`.

```js
assert.strictEqual(document.querySelector('.table-wrap')?.localName, 'div');
```

Ваш елемент `.table-wrap` має йти після вашого існуючого елемента `div`.

```js
assert.isTrue(
  document.querySelectorAll('div')?.[1]?.classList?.contains('table-wrap')
);
```

Ваш елемент `.table-wrap` має містити три елементи `table`.

```js
const children = [...(document.querySelector('.table-wrap')?.children ?? [])];
assert.lengthOf(children, 3);
children.forEach(child => assert.strictEqual(child?.localName, 'table'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
