---
id: 61fd6343fb42013d99bcd7f3
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Усередині вашого `body` елемента вкладіть `section` елемент всередину `main` елемента.

# --hints--

Першим елементом всередині `body` елемента має бути ваш `main` елемент.

```js
assert.strictEqual(
  document.querySelector('body')?.children?.[0]?.localName,
  'main'
);
```

У вашому `section` елементі має бути `main` елемент.

```js
assert.strictEqual(
  document.querySelector('main')?.children?.[0]?.localName,
  'section'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```

```css

```
