---
id: 61fd7a160ed17960e971f28b
title: Крок 16
challengeType: 0
dashedName: step-16
---

# --description--

У вашому четвертому `tr` елементі додайте `th` елемент із текстом `Total Assets`. Обгорніть текст `Assets` у `span` елемент із атрибутом `class`, встановленим у `sr-only`.

Після цього додайте три `td` елементи з таким текстом (у порядку): `$579`, `$736`, `$809`. Дайте третьому `td` елементу атрибут `class`, встановлений у `current`.

# --hints--

Ваш четвертий `tr` має містити `th` елемент.

```js
assert.exists(
  document
    .querySelector('tbody')
    ?.querySelectorAll('tr')?.[3]
    ?.querySelector('th')
);
```

Ваш `th` елемент має містити текст `Total Assets`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableHeader = tableRow.querySelector('th');
assert.strictEqual(tableHeader.innerText.trim(), 'Total Assets');
```

Ви маєте обгорнути текст `Assets` у `span` елемент.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.strictEqual(span.textContent.trim(), 'Assets');
```

Ваш `span` елемент має мати атрибут `class`, встановлений у `sr-only`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const span = tableRow.querySelector('th > span');
assert.isTrue(span?.classList?.contains('sr-only'));
```

Ви маєте мати три `td` елементи.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableDatas = tableRow.querySelectorAll('td');
assert.lengthOf(tableDatas, 3);
```

Ваш перший `td` елемент має містити текст `$579`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[0];
assert.strictEqual(tableData?.textContent.trim(), '$579');
```

Ваш другий `td` елемент має містити текст `$736`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[1];
assert.strictEqual(tableData?.textContent.trim(), '$736');
```

Ваш третій `td` елемент має містити текст `$809`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.strictEqual(tableData.textContent.trim(), '$809');
```

Ваш третій `td` елемент має мати атрибут `class`, встановлений у `current`.

```js
const tableRow = document.querySelector('tbody')?.querySelectorAll('tr')?.[3];
const tableData = tableRow.querySelectorAll('td')?.[2];
assert.isTrue(tableData?.classList?.contains('current'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
              <tr class="data">
                <th>Checking <span class="description">Our primary transactional account.</span></th>
                <td>$54</td>
                <td>$56</td>
                <td class="current">$53</td>
              </tr>
              <tr class="data">
                <th>Savings <span class="description">Funds set aside for emergencies.</span></th>
                <td>$500</td>
                <td>$650</td>
                <td class="current">$728</td>
              </tr>
--fcc-editable-region--
              <tr class="total">
              </tr>
--fcc-editable-region--
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
