---
id: 61fd5a93fd62bb35968adeab
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Почніть свій проєкт, додавши до вашого `head` елемента `link` елемент для вашого стилю.

# --hints--

Ваш код має містити `link` елемент.

```js
assert.match(code, /<link/);
```

Ваш `link` елемент має бути всередині вашого `head` елемента.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

Ваш `link` елемент має мати атрибут `rel` зі значенням `stylesheet`.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

Ваш `link` елемент має мати атрибут `href` зі значенням `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
