---
id: 61fd66c687e610436494c6f1
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У вашому `section` елементі додайте `h1` елемент з вкладеним `span` елементом.

# --hints--

Ваш `section` елемент має містити `h1` елемент.

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

Ваш `h1` елемент має містити `span` елемент.

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
