---
id: 63ec20a06fff670d37befbd9
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이제 `Dessert` 클래스가 있으니 이를 사용해 제품을 만들 수 있습니다. `new Dessert(...)`를 호출할 때마다 같은 형태의 새 인스턴스가 반환되며, 전체 객체 리터럴을 작성하는 것보다 훨씬 간단합니다.

`products` 배열에 단일 값을 추가하세요. `new Dessert()`를 호출할 때 인수로 숫자 `1`, 문자열 `"Vanilla Cupcakes (6 Pack)"`, 숫자 `12.99`, 문자열 `"Cupcake"`를 순서대로 전달하세요.

# --hints--

`products` 배열에는 값이 하나 있어야 합니다.

```js
assert.strictEqual(products.length, 1);
```

`products` 배열의 첫 번째 값은 `Dessert`의 인스턴스여야 합니다.

```js
assert.instanceOf(products[0], Dessert);
```

첫 번째 제품은 `id` 속성이 숫자 `1`로 설정되어 있어야 합니다.

```js
assert.strictEqual(products[0].id, 1);
```

첫 번째 제품은 `name` 속성이 문자열 `"Vanilla Cupcakes (6 Pack)"`로 설정되어 있어야 합니다.

```js
assert.strictEqual(products[0].name, 'Vanilla Cupcakes (6 Pack)');
```

첫 번째 제품은 `price` 속성이 숫자 `12.99`로 설정되어 있어야 합니다.

```js
assert.strictEqual(products[0].price, 12.99);
```

첫 번째 제품은 `category` 속성이 문자열 `"Cupcake"`로 설정되어 있어야 합니다.

```js
assert.strictEqual(products[0].category, 'Cupcake');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp Shopping Cart</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1 class="title">Desserts Page</h1>
    </header>
    <main>
      <button id="cart-btn" type="button" class="btn">
        <span id="show-hide-cart">Show</span> Cart
      </button>
      <div id="cart-container">
        <button class="btn" id="clear-cart-btn">Clear Cart</button>
        <div id="products-container"></div>
        <p>Total number of items: <span id="total-items">0</span></p>
        <p>Subtotal: <span id="subtotal">$0</span></p>
        <p>Taxes: <span id="taxes">$0</span></p>
        <p>Total: <span id="total">$0</span></p>
      </div>
      <div id="dessert-card-container"></div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.title {
  color: var(--light-grey);
  text-align: center;
  margin: 25px 0;
}

#dessert-card-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
}

.dessert-card {
  background-color: var(--light-grey);
  padding: 15px;
  text-align: center;
  border-radius: 15px;
  margin: 20px 10px;
}

.dessert-price {
  font-size: 1.2rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 100px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#cart-btn {
  position: absolute;
  top: 0;
  right: 0;
}

.add-to-cart-btn {
  margin: 30px auto 10px;
}

#cart-container {
  display: none;
  position: absolute;
  top: 60px;
  right: 0;
  background-color: var(--light-grey);
  width: 200px;
  height: 400px;
  border: 8px double var(--black);
  border-radius: 15px;
  text-align: center;
  font-size: 1.2rem;
  overflow-y: scroll;
}

.product {
  margin: 25px 0;
}

.product-count {
  display: inline-block;
  margin-right: 10px;
}

.product-category {
  margin: 10px 0;
}

@media (min-width: 768px) {
  #dessert-card-container {
    flex-direction: row;
  }

  .dessert-card {
    flex: 1 0 21%;
  }

  #cart-container {
    width: 300px;
  }
}
```

```js
const cartContainer = document.getElementById("cart-container");
const productsContainer = document.getElementById("products-container");
const dessertCards = document.getElementById("dessert-card-container");
const cartBtn = document.getElementById("cart-btn");
const clearCartBtn = document.getElementById("clear-cart-btn");
const totalNumberOfItems = document.getElementById("total-items");
const cartSubTotal = document.getElementById("subtotal");
const cartTaxes = document.getElementById("taxes");
const cartTotal = document.getElementById("total");
const showHideCartSpan = document.getElementById("show-hide-cart");
let isCartShowing = false;

class Dessert {
  constructor(id, name, price, category) {
    this.id = id;
    this.name = name;
    this.price = price;
    this.category = category;
  }
}

const products = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
