---
id: 69d7d4dc8c6e21a2a19c4d34
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

HTML `class`에 이미 익숙하지만, JavaScript에도 <dfn>클래스</dfn>가 있습니다. JavaScript에서 클래스는 객체를 생성하기 위한 청사진과 같습니다. 속성과 메서드 집합을 정의하고, 그 속성과 메서드를 가진 새 객체를 인스턴스화(또는 생성)할 수 있게 해줍니다.

다음은 `Dessert` 클래스를 선언하는 예입니다:

```js
class Dessert {}
```

제품을 추가하기 전에 `class`라는 `Dessert`를 만드세요. 네 개의 매개변수를 순서대로 받는 `constructor`를 만드세요: `id`, `name`, `price`, 그리고 `category`. `constructor` 안에서 각 매개변수를 `this`의 같은 이름을 가진 속성에 할당하세요.

# --hints--

`Dessert`라는 이름의 클래스를 선언해야 합니다.

```js
assert.isFunction(Dessert);
```

`Dessert` 클래스에는 네 개의 매개변수 `constructor`, `id`, `name`, 그리고 `price`를 받는 `category`가 있어야 합니다.

```js
assert.strictEqual(Dessert.length, 4);
```

`Dessert` 생성자는 각 매개변수를 `this`의 같은 이름을 가진 속성에 할당해야 합니다.

```js
const d = new Dessert(1, 'Test', 1.99, 'Test Category');
assert.strictEqual(d.id, 1);
assert.strictEqual(d.name, 'Test');
assert.strictEqual(d.price, 1.99);
assert.strictEqual(d.category, 'Test Category');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp Shopping Cart</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <header>
      <h1 class="title">Desserts Page</h1>
    </header>
    <main>
      <button id="cart-btn" type="button" class="btn">
        <span id="show-hide-cart">Show</span> Cart
      </button>
      <div id="cart-container">
        <button class="btn" id="clear-cart-btn">Clear Cart</button>
        <div id="products-container"></div>
        <p>Total number of items: <span id="total-items">0</span></p>
        <p>Subtotal: <span id="subtotal">$0</span></p>
        <p>Taxes: <span id="taxes">$0</span></p>
        <p>Total: <span id="total">$0</span></p>
      </div>
      <div id="dessert-card-container"></div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #1b1b32;
  --light-grey: #f5f6f7;
  --black: #000;
  --white: #fff;
  --grey: #3b3b4f;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  background-color: var(--dark-grey);
}

.title {
  color: var(--light-grey);
  text-align: center;
  margin: 25px 0;
}

#dessert-card-container {
  display: flex;
  flex-direction: column;
  flex-wrap: wrap;
  align-items: center;
}

.dessert-card {
  background-color: var(--light-grey);
  padding: 15px;
  text-align: center;
  border-radius: 15px;
  margin: 20px 10px;
}

.dessert-price {
  font-size: 1.2rem;
}

.btn {
  display: block;
  cursor: pointer;
  width: 100px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#cart-btn {
  position: absolute;
  top: 0;
  right: 0;
}

.add-to-cart-btn {
  margin: 30px auto 10px;
}

#cart-container {
  display: none;
  position: absolute;
  top: 60px;
  right: 0;
  background-color: var(--light-grey);
  width: 200px;
  height: 400px;
  border: 8px double var(--black);
  border-radius: 15px;
  text-align: center;
  font-size: 1.2rem;
  overflow-y: scroll;
}

.product {
  margin: 25px 0;
}

.product-count {
  display: inline-block;
  margin-right: 10px;
}

.product-category {
  margin: 10px 0;
}

@media (min-width: 768px) {
  #dessert-card-container {
    flex-direction: row;
  }

  .dessert-card {
    flex: 1 0 21%;
  }

  #cart-container {
    width: 300px;
  }
}
```

```js
const cartContainer = document.getElementById("cart-container");
const productsContainer = document.getElementById("products-container");
const dessertCards = document.getElementById("dessert-card-container");
const cartBtn = document.getElementById("cart-btn");
const clearCartBtn = document.getElementById("clear-cart-btn");
const totalNumberOfItems = document.getElementById("total-items");
const cartSubTotal = document.getElementById("subtotal");
const cartTaxes = document.getElementById("taxes");
const cartTotal = document.getElementById("total");
const showHideCartSpan = document.getElementById("show-hide-cart");
let isCartShowing = false;

--fcc-editable-region--

--fcc-editable-region--

const products = [];
```
