---
id: 694bf6eb530e19a9c48c59f7
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Ви також можете використовувати від’ємні числа, щоб вирізати частину рядка з кінця. Наприклад, `-1` позначає останній символ, `-2` — передостанній символ і так далі.

Щоб отримати останні три символи рядка, можна використати `string[-3:]`. Зверніть увагу, що параметр stop після двокрапки пропущено. Це означає, що вирізання має тривати до межі рядка.

Створіть змінну з ім’ям `last_three`. Використайте від’ємну індексацію, щоб витягти останні три символи з `employee_code` (які представляють ідентифікатор). Нарешті, виведіть `last_three` у термінал.

На цьому практичне заняття зі створення картки працівника завершене.

# --hints--

Ви маєте мати змінну з ім’ям `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

Ви маєте вирізати останні три символи з `employee_code` за допомогою від’ємної індексації і призначити їх змінній `last_three`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

Ви маєте вивести `last_three` у термінал.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
