---
id: 694bf6eb530e19a9c48c59f6
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Ви можете вирізати частину рядка з будь-якої позиції, а не лише з початку. Це корисно у багатьох випадках.

Створіть змінну `year_code` і призначте їй зріз `employee_code` від індексу `4` до `8`. Це дозволить витягти `2026`.

Потім створіть змінну `initials` і призначте їй зріз `employee_code` від індексу `9` до `11`. Це дозволить витягти `JD`.

Нарешті, виведіть обидві змінні в термінал.

# --hints--

Ви повинні мати змінну з ім’ям `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

Ви повинні вирізати зріз `employee_code` від індексу `4` до `8` і призначити результат змінній `year_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

Ви повинні мати змінну з ім’ям `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

Ви повинні вирізати зріз `employee_code` від індексу `9` до `11` і призначити результат змінній `initials`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

Ви повинні вивести `year_code` у термінал.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

Ви повинні вивести `initials` у термінал.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
