---
id: 694bf6eb530e19a9c48c59f1
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер ваша адреса здається неповною. Ви також хочете додати номер квартири, де живе працівник, тому слід змінити змінну.

Коли потрібно додати вміст у кінець існуючої рядкової змінної, можна використати **оператор розширеного присвоєння** ,`+=`. Це коротше, ніж писати `var = var + 'new text'`. Наприклад:

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

Пам’ятайте, що рядки є незмінними, тому ця операція не змінює оригінальний рядок. Замість цього створюється новий рядок і призначається змінній.

Використайте оператор `+=`, щоб додати рядок `, Apartment 4B` до вашої змінної `address`.

# --hints--

Ви повинні використати оператор `+=`, щоб додати рядок `, Apartment 4B` до змінної `address`.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
