---
id: 694bf3af05d10fd89dab05cc
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

У Python ви можете об’єднувати рядки за допомогою оператора `+`. Це називається конкатенацією. Ось приклад:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

Створіть змінну `full_name`, об’єднавши `first_name` та `last_name`. Потім виведіть `full_name`.

# --hints--

У вас має бути змінна з іменем `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

Ви маєте призначити змінній `first_name` рядок, утворений конкатенацією `last_name` та `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

Ви маєте вивести змінну `full_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
