---
id: 694bf6eb530e19a9c48c59f5
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Працюючи з рядками, вам часто потрібно буде витягти певну частину рядка. Це називається **вирізанням**.

Синтаксис такий `string[start:stop]`, де:

* `start` — індекс, з якого починається вирізання (**включно**).
* `stop` — індекс, на якому вирізання закінчується (**виключно**).

Наприклад, якщо `text = 'Python'`, тоді `text[0:2]` дає `'Py'`.

Визначте `employee_code` як `'DEV-2026-JD-001'`. Після цього створіть змінну `department` і призначте їй вирізку з `employee_code` від індексу `0` до `3`. Потім виведіть `department` у термінал.

# --hints--
У вас має бути змінна з ім’ям `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

Ви маєте призначити рядок `DEV-2026-JD-001` змінній `employee_code`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

У вас має бути змінна з ім’ям `department`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

Ви маєте вирізати перші три символи з `employee_code` і призначити їх змінній `department`.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

Ви маєте вивести `department` у термінал.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
