---
id: 694bf6eb530e19a9c48c59f4
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Конкатенація багатьох рядків за допомогою `+` і перетворення чисел за допомогою `str()` може стати заплутаною і важкою для читання.

У Python 3.6 з’явилися **f-рядки** для вирішення цієї проблеми. Додавши літеру `f` перед відкриваючою лапкою, ви можете вставляти змінні та вирази всередину полів заміни, які позначаються фігурними дужками `{}`. Наприклад:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

Створіть змінну `employee_card` і призначте їй f-рядок, який відображає `Employee:`, за яким іде пробіл і значення змінної `full_name`.

# --hints--

У вас має бути змінна з іменем `employee_card`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

Значення вашої змінної `employee_card` має бути `f'Employee: {full_name}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
