---
id: 694bf6ea530e19a9c48c59f0
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Якщо ви конкатенуєте два рядки, як `'John' + 'Doe'`, результатом буде `'JohnDoe'` без пробілу. Щоб це виправити, потрібно конкатенувати рядок, що містить пробіл (`' '`), між ними.

Оновіть свою змінну `full_name`, щоб вона конкатенувала `first_name`, пробіл і `last_name`.

# --hints--

Ви повинні додати пробіл (`' '`) між `first_name` і `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
