---
id: 694bf6eb530e19a9c48c59f3
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер ви знову використаєте функцію `str()`. Як і з age, потрібно перетворити будь-яку числову змінну на рядок перед тим, як об’єднувати її з іншим текстом.

Створіть змінну з іменем `experience_years` і призначте їй ціле число `5`.

Потім створіть змінну `experience_info`. Призначте їй рядок, утворений шляхом конкатенації `'Experience: '`, змінної `experience_years` (перетвореної на рядок) та `' years'`. Виведіть результат у термінал.

# --hints--

У вас має бути змінна з іменем `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

Ви повинні призначити ціле число `5` змінній `experience_years`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

У вас має бути змінна з іменем `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

Ви повинні призначити рядок, утворений конкатенацією `'Experience: '`, `str(experience_years)` та `' years'`, змінній `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

Виведіть `experience_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
