---
id: 695c96f82af4fb78e59fa9ed
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Рядки — це послідовності символів, які використовують для збереження текстових даних. Як ви, можливо, пам’ятаєте з попередніх уроків, рядок можна створити, взявши текст у одинарні (`'`) або подвійні (`"`) лапки. Наприклад:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

Створіть змінну `first_name`, яка зберігає рядок `John`, і змінну `last_name`, яка зберігає рядок `Doe`. Потім виведіть `first_name` і `last_name`.

# --hints--

У вас має бути змінна з ім’ям `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

Ви маєте призначити рядок `John` змінній `first_name`. Не забудьте взяти текст у одинарні або подвійні лапки.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

У вас має бути змінна з ім’ям `last_name`.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

Ви маєте призначити рядок `Doe` змінній `last_name`. Не забудьте взяти текст у одинарні або подвійні лапки.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

Виведіть у термінал змінну `first_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

Виведіть у термінал змінну `last_name`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
