---
id: 695c9f21a8b1d53cf1b97841
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер створіть змінну з іменем `employee_age` і призначте їй ціле число `28`. 


# --hints--

У вас має бути змінна з іменем `employee_age`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

Ви повинні призначити ціле число `28` вашій змінній `employee_age`. Не беріть текст у лапки.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
