---
id: 695c9c7c61d446f1451e0ae7
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Далі створіть змінну `address` для збереження адреси працівника. Призначте їй рядок `123 Main Street`, а потім виведіть `address`.

# --hints--

У вас має бути змінна з іменем `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

Ви повинні призначити рядок `123 Main Street` змінній `address`. Не забудьте взяти текст у одинарні або подвійні лапки.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

Ви повинні вивести змінну `address`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
