---
id: 695df6480de23758af5a29cd
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Тепер ви хочете створити рядок, який відображає вік працівника.

Почніть зі створення змінної `employee_info` і призначте їй результат конкатенації:

- змінної `full_name`.
- рядка, що складається з символів `is`, перед якими і після яких стоїть пробіл.

# --hints--

У вас має бути змінна з ім’ям `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

Ви повинні призначити `full_name + ' is '` вашій змінній `employee_info`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
