---
id: 695ea1e55c0e619e597fc5c9
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Як ви бачите, Python викликав `TypeError: can only concatenate str (not "int") to str`. Це відбувається тому, що Python не дозволяє безпосередньо конкатенувати текст (рядки) і числа (цілі числа).

Щоб виправити це, потрібно спочатку перетворити число на рядок за допомогою функції `str()`, яка повертає рядкову версію об’єкта:

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

Оновіть своє присвоєння `employee_info`, щоб перетворити `employee_age` на рядок за допомогою `str(employee_age)`.

# --hints--

Вам слід змінити присвоєння змінної `employee_info`, перетворивши `employee_age` на рядок за допомогою `str(employee_age)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
