---
id: 695fd3dde03273875e5d6709
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Оновіть присвоєння `employee_info`, щоб також додати `employee_age` у кінці.

Після цього ви побачите `TypeError` у терміналі. У наступному кроці ви працюватимете над його виправленням.

# --hints--

Вам слід оновити присвоєння `employee_info`, щоб додати `employee_age` у кінці.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
