---
id: 69727ea5c9606bedfe36d3d2
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Тепер час додати фінальні деталі до картки.

Створіть змінну з іменем `position` зі значенням рядка `Data Analyst` і змінну з іменем `salary` зі значенням цілого числа `75000`.

Потім оновіть свій f-рядок `employee_card`, щоб включити посаду та зарплату. Він має мати саме такий формат: `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`. Замініть заповнювачі відповідними змінними.

Нарешті, виведіть `employee_card`, щоб побачити результат.

# --hints--

У вас має бути змінна з іменем `position`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

Ви маєте призначити рядок `Data Analyst` змінній `position`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

У вас має бути змінна з іменем `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

Ви маєте призначити ціле число `75000` змінній `salary`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

Ви маєте оновити `employee_card`, щоб використовувати f-рядок `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

Ви маєте вивести `employee_card` у термінал.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
