---
id: 6980d3ecfe800d8a54130f30
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Наразі `employee_card` показує лише ім’я працівника. Тепер ви додасте більше інформації.

Оновіть присвоєння `employee_card`, щоб включити вік працівника. Фінальний рядок має виглядати так: `Employee: [name] | Age: [age]` з `[name]`, заміненим на ім’я працівника, і `[age]`, заміненим на вік працівника.

# --hints--

Ви маєте оновити `employee_card`, щоб використовувати f-рядок `f'Employee: {full_name} | Age: {employee_age}'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
