---
id: afcc8d540bea9ea2669306b6
title: Eine Funktion zum Wiederholen eines Strings erstellen
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die einen gegebenen String eine bestimmte Anzahl von Malen wiederholt. Für dieses Labor verwenden Sie _nicht_ die eingebaute `.repeat()`-Methode.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `repeatStringNumTimes` erstellen, die zwei Parameter entgegennimmt: einen String und eine Zahl.
2. Die Funktion sollte den String so oft wiederholen, wie es die angegebene Zahl vorgibt.
3. Wenn die Zahl kleiner oder gleich null ist, sollte die Funktion einen leeren String zurückgeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `repeatStringNumTimes` erstellen.

```js
assert.isFunction(repeatStringNumTimes);
```

`repeatStringNumTimes` sollte zwei Parameter entgegennehmen.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

Die Funktion `repeatStringNumTimes` sollte immer einen String zurückgeben.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

`repeatStringNumTimes("*", 3)` sollte den String `***` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

`repeatStringNumTimes("abc", 3)` sollte den String `abcabcabc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

`repeatStringNumTimes("abc", 4)` sollte den String `abcabcabcabc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

`repeatStringNumTimes("abc", 1)` sollte den String `abc` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

`repeatStringNumTimes("*", 8)` sollte den String `********` zurückgeben.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

`repeatStringNumTimes("abc", -2)` sollte einen leeren String (`""`) zurückgeben.

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

`repeatStringNumTimes("abc", 0)` sollte `""` zurückgeben.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

Die eingebaute `repeat()`-Methode sollte nicht verwendet werden.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
