---
id: 686e64be3cd9db25282a956d
title: Tengeneza muundo wa maonyesho ya upigaji picha
challengeType: 25
dashedName: lab-photography-exhibit
demoType: onClick
---

# --description--

Katika maabara hii, utaunda muundo wa maonyesho ya upigaji picha kwa kutumia madarasa ya huduma ya Tailwind CSS.

Kwa picha, unaweza kutumia URL zifuatazo ikiwa unataka:

- https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg
- https://cdn.freecodecamp.org/curriculum/labs/alps.jpg
- https://cdn.freecodecamp.org/curriculum/labs/sea.jpg

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kipengele chenye darasa la `main-container`.
2. Kipengele chako cha `.main-container` kinapaswa pia kuwa na madarasa ya huduma ya Tailwind CSS yafuatayo:
   - darasa la huduma sahihi la kuunda chombo cha gridi ya CSS.
   - darasa la huduma la kufafanua idadi ya nguzo ndani ya gridi. Unapaswa kutumia darasa la huduma linalotumia idadi thabiti ya nguzo. Mfano `grid-cols-<number>`.
   - darasa la huduma la kuweka nafasi ya safu na nguzo kwa vipande vya gridi.
3. Ndani ya kipengele chako cha `.main-container`, kunapaswa kuwa na vipengele angalau vitatu kila kikiwa na darasa la `card`.
4. Kila kipengele chako cha `.card` kinapaswa pia kuwa na madarasa ya huduma ya Tailwind CSS yafuatayo:
   - darasa la huduma la kuweka mpaka uliozungushwa kwa chaguo lako.
   - darasa la huduma la kuweka nafasi ya ndani kwa chaguo lako. Unapaswa kutumia muundo huu `p-<number>`.
5. Ndani ya kila kipengele chako cha `.card`, unapaswa kuwa na vipengele vifuatavyo:
   - kipengele cha picha chenye sifa za `src` na `alt`.
   - kipengele chenye darasa la `subheading`.
   - kipengele chenye darasa la `description`.
6. Kila kipengele chako cha picha kinapaswa kuwa na darasa la huduma la kuweka mpaka uliozungushwa kwa chaguo lako.
7. Kila kipengele chako cha `.subheading` kinapaswa kutumia madarasa ya huduma kuweka uzito na ukubwa wa fonti kwa chaguo lako.
8. Kila kipengele chako cha `.description` kinapaswa kutumia darasa la huduma kuweka ukubwa wa fonti kwa chaguo lako.

# --hints--

Unapaswa kuwa na kipengele chenye darasa la `main-container`.

```js
assert.exists(document.querySelector(".main-container"));
```

Kipengele chako cha `.main-container` kinapaswa kutumia darasa sahihi la huduma ya Tailwind CSS kuunda chombo cha gridi ya CSS.

```js
const mainContainer = document.querySelector(".main-container");
assert.isTrue(mainContainer.classList.contains("grid"));
```

Kipengele chako cha `.main-container` kinapaswa kutumia darasa la huduma la kufafanua idadi ya nguzo ndani ya gridi.

```js
const mainContainer = document.querySelector(".main-container");
// Grid classes can vary, so this check is for the presence of any grid-cols class
const hasGridCols = [...mainContainer.classList].some(className => className.startsWith("grid-cols-"));
assert.isTrue(hasGridCols);
```

Kipengele chako cha `.main-container` kinapaswa kutumia darasa la huduma la kuweka nafasi ya safu na nguzo kwa vipande vya gridi.

```js
const mainContainer = document.querySelector(".main-container");
// Check for any gap class including gap-x and gap-y

const hasGap = [...mainContainer.classList].some(c => c.startsWith("gap-")) ||
  (classList.some(c => c.startsWith("gap-x-")) &&
   classList.some(c => c.startsWith("gap-y-")));

assert.isTrue(hasGap);
```

Kipengele chako cha `.main-container` kinapaswa kuwa na vipengele angalau vitatu kila kikiwa na darasa la `card`.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isAtLeast(cards.length, 3);
```

Kila kipengele chako cha `.card` kinapaswa kuwa na darasa la huduma la kuweka mpaka uliozungushwa.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);

const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

cards.forEach(card => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => card.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
})
```

Kila kipengele chako cha `.card` kinapaswa kuwa na darasa la huduma la kuweka nafasi ya ndani.

```js
const cards = document.querySelectorAll(".main-container .card");
assert.isNotEmpty(cards);
const hasPadding = [...cards].every(card => {
  return [...card.classList].some(c => c.startsWith("p-"));
});
assert.isTrue(hasPadding);
```

Kila kipengele chako cha `.card` kinapaswa kuwa na kipengele cha picha.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
```

Kila kipengele chako cha `img` kinapaswa kuwa na sifa ya `src`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("src")));
images.forEach(img => assert.isTrue(img.getAttribute("src").length > 0));
```

Kila kipengele chako cha `img` kinapaswa kuwa na sifa ya `alt`.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isAtLeast(images.length, 3);
images.forEach(img => assert.isTrue(img.hasAttribute("alt")));
images.forEach(img => assert.isTrue(img.getAttribute("alt").length > 0));
```

Kila kipengele chako cha `.card` kinapaswa kuwa na kipengele chenye darasa la `subheading`.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isAtLeast(subheadings.length, 3);
```

Kila kipengele chako cha `.subheading` kinapaswa kuwa na maudhui ya maandishi.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(subheadings);

subheadings.forEach(subheading => {
  assert.isNotEmpty(subheading.textContent.trim());
});
```

Kila kipengele chako cha `.card` kinapaswa kuwa na kipengele chenye darasa la `description`.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
assert.isAtLeast(descriptions.length, 3);
```

Kila kipengele chako cha `.description` kinapaswa kuwa na maudhui ya maandishi.

```js
const descriptions = document.querySelectorAll(".main-container .card .description");
// Ensure we are not dealing with an empty NodeList
assert.isNotEmpty(descriptions);

descriptions.forEach(description => {
  assert.isNotEmpty(description.textContent.trim());
});
```

Kila kipengele chako cha `img` kinapaswa kuwa na darasa la huduma la kuweka mpaka uliozungushwa.

```js
const images = document.querySelectorAll(".main-container .card img");
assert.isNotEmpty(images);
const predefinedRoundedCorners = [
  "rounded-none",
  "rounded-sm",
  "rounded",
  "rounded-md",
  "rounded-lg",
  "rounded-xl",
  "rounded-2xl",
  "rounded-3xl",
  "rounded-full"
];

images.forEach(img => {
  const hasRoundedCorners = predefinedRoundedCorners.some(corner => img.classList.contains(corner));
  assert.isTrue(hasRoundedCorners);
});
```

Kila kipengele chako cha `.subheading` kinapaswa kutumia madarasa ya huduma kuweka uzito wa fonti.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedFontWeights = [
  "font-thin",
  "font-extralight",
  "font-light",
  "font-normal",
  "font-medium",
  "font-semibold",
  "font-bold",
  "font-extrabold",
  "font-black"
];

subheadings.forEach(subheading => {
  const hasFontWeight = predefinedFontWeights.some(weight => subheading.classList.contains(weight));
  assert.isTrue(hasFontWeight);
});
```

Kila kipengele chako cha `.subheading` kinapaswa kutumia madarasa ya huduma kuweka ukubwa wa fonti.

```js
const subheadings = document.querySelectorAll(".main-container .card .subheading");
assert.isNotEmpty(subheadings);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

subheadings.forEach(subheading => {
  const hasFontSize = predefinedSizes.some(size => subheading.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

Kila kipengele chako cha `.description` kinapaswa kutumia darasa la huduma kuweka ukubwa wa fonti.

```js
const descriptions = document.querySelectorAll(".description");
assert.isNotEmpty(descriptions);
const predefinedSizes = [
  "text-xs",
  "text-sm",
  "text-base",
  "text-lg",
  "text-xl",
  "text-2xl",
  "text-3xl",
  "text-4xl",
  "text-5xl",
  "text-6xl",
  "text-7xl",
  "text-8xl",
  "text-9xl"
];

descriptions.forEach(description => {
  const hasFontSize = predefinedSizes.some(size => description.classList.contains(size));
  assert.isTrue(hasFontSize);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
 
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Photography Exhibit</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-purple-800 text-yellow-400 min-h-screen">
    <h1 class="text-center text-4xl font-bold my-12">Photography Exhibit</h1>

    <main class="main-container grid grid-cols-1 md:grid-cols-3 gap-8 px-4">
      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="Ancient Colosseum"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Colosseum, Rome</h2>
          <p class="description text-sm mb-4">
            There is so much history and beauty. The light hits it just right.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="Alpine Mountains"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">The Alps</h2>
          <p class="description text-sm mb-4">
            The mountains go on forever and the view is breathtaking.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>

      <figure
        class="card bg-purple-900 border border-yellow-400 rounded-lg p-4"
      >
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="Calm Sea"
          class="h-48 w-full object-cover rounded-md mb-4"
        />
        <figcaption>
          <h2 class="subheading text-lg font-semibold mb-2">Seascape</h2>
          <p class="description text-sm mb-4">
            The water was still, the air was cool and the feeling is pure peace.
          </p>
          <div class="grid grid-cols-2 gap-4">
            <button
              class="bg-purple-800 hover:bg-purple-700 border border-yellow-400 py-1 px-2 rounded-md"
            >
              Like
            </button>
            <a
              href="#"
              class="bg-yellow-400 hover:bg-yellow-300 text-purple-900 font-semibold py-1 px-2 rounded-md text-center"
            >
              View Details
            </a>
          </div>
        </figcaption>
      </figure>
    </main>
  </body>
</html>
```
