---
id: 683ec95f2f8781355556ff80
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Lo siguiente que debe hacer es crear una condición que verifique si el `value` objetivo está en el medio. 

Dentro del bucle `while`, cree una declaración `if` que verifique si el `value` objetivo es igual a `value_at_middle`. Si lo es, devuelva la variable `path_to_target`.

# --hints--

Debe crear una declaración `if` que verifique si el `value` es igual a `value_at_middle`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_conditions()[0]
assert _cond.is_equivalent("value == value_at_middle") or _cond.is_equivalent("value_at_middle == value")
`) })
```

Debe devolver la variable `path_to_target` dentro de la declaración `if`.

```js
({
    test: () =>
      assert(
        runPython(
          `_Node(_code).find_function("binary_search").find_whiles()[0].find_body().find_ifs()[0].find_body().is_equivalent("return path_to_target")`
        )
      )
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)
--fcc-editable-region--
        
--fcc-editable-region--
```
