---
id: 683ec95f2f8781355556ff7a
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

El algoritmo de búsqueda binaria que estás construyendo no solo retornará `True` o `False` para indicar si el valor fue encontrado o no. También devolverá el camino que toma para encontrar ese valor. Y para eso, tendrás una variable `path_to_target` para seguir el camino.

Dentro de tu función `binary_search`, reemplaza la palabra clave `pass` con una variable `path_to_target` e inicialízala como una lista vacía, para comenzar.

# --hints--

Deberías crear una variable `path_to_target` que se establezca como una lista vacía dentro de tu función `binary_search`. No olvides eliminar la palabra clave `pass`.

```js
({ 
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_body().is_equivalent("path_to_target = []") and not _Node(_code).find_function("binary_search").has_pass()`
      )
    ) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    pass
--fcc-editable-region--
```
