---
id: 683ec95f2f8781355556ff84
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Si la condición en el `elif` es cierta, entonces quite la palabra clave `pass` y actualice el valor de la variable `low` sumando `1` a la variable `mid`.

Esto ampliará la búsqueda a la mitad derecha de las áreas de búsqueda actuales en la lista, porque si el `value` es mayor que `value_at_middle`, significa que el `value` debe estar en la mitad derecha de la zona de búsqueda actual.

# --hints--

Deberías quitar la palabra clave `pass` y actualizar la variable `low` a `mid + 1`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).find_function("binary_search").find_whiles()[0].find_ifs()[0].find_bodies()[1].is_equivalent("low = mid + 1")`))
})
```

# --seed--

## --seed-contents--

```py
def binary_search(search_list, value):
    path_to_target = []
    low = 0
    high = len(search_list) - 1
    
    while low <= high:
        mid = (low + high) // 2
        value_at_middle = search_list[mid]
        path_to_target.append(value_at_middle)

        if value == value_at_middle:
            return path_to_target
--fcc-editable-region--
        elif value > value_at_middle:
            pass
--fcc-editable-region--
        break
    return []

print(binary_search([1, 2, 3, 4, 5], 3))
print(binary_search([1, 2, 3, 4, 5, 9], 4)) 
```
