---
id: 683ec95f2f8781355556ff7b
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En la búsqueda binaria, hay el índice más bajo posible y el índice más alto posible.

El índice más bajo posible representa el límite más a la izquierda en la parte actual de los valores que se están buscando en, en este caso, `search_list`. El índice más alto posible, por otro lado, representa el límite más a la derecha en los valores que se están buscando.

Para tener en cuenta esos dos, debes considerar toda la lista que se está buscando. Entonces, dentro de la función `binary_search`, define una variable `low` con un valor de `0`, y una variable `high` con un valor que tenga en cuenta el último índice de la lista que se está buscando.

# --hints--

Debes tener una variable llamada `low` en tu función `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("low")`)) 
})
```

La variable `low` debe establecerse en `0`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_function("binary_search").find_variable("low").is_equivalent("low = 0")`
      )
    )
})
```

Debes tener una variable llamada `high` en tu función `binary_search`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("binary_search").has_variable("high")`)) 
})
```

La variable `high` debe establecerse en `len(search_list) - 1`.

```js
({
  test: () => assert(
    runPython(
    `_Node(_code).find_function("binary_search").find_variable("high").is_equivalent("high = len(search_list) - 1")`
    ))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def binary_search(search_list, value):
    path_to_target = []

--fcc-editable-region--
```
