---
id: ab306dbdcc907c7ddfc30830
title: Tengeneza chombo cha kupangilia safu ya data kwa kina
challengeType: 26
dashedName: create-a-deep-flattening-tool
---

# --description--

Katika maabara hii, utatekeleza algoriti ya kupangilia safu ya data.

Kupangilia safu ya data kunamaanisha kubadilisha safu ya data zilizopangwa ndani kwa kina chochote kuwa safu moja ya kipimo kimoja. Mchakato huu huchukua vipengele vyote kwa mpangilio, ukifungua tu safu za data. Aina nyingine hazibadiliki.

Kwa mfano:

| Asili | Iliyopangiliwa |
| --- | --- |
| `[[1], [], [2, [3]]]` | `[1, 2, 3]` |
| `[1, {"foo": "bar"}, [2]]` | `[1, {"foo": "bar"}, 2]` |
| `["baz", [1, 2], {}]` | `["baz", 1, 2, {}]` |

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuwa na kitendakazi kinachoitwa `steamrollArray`.
1. Kitendakazi `steamrollArray` kinapaswa kupokea hoja moja: safu ya data iliyopangwa ndani.
1. Kitendakazi kinapaswa kupangilia safu ya data iliyopangwa ndani, kikizingatia viwango tofauti vya upangiliaji ndani.
1. Suluhisho lako halipaswi kutumia njia za `Array.prototype.flat()` au `Array.prototype.flatMap()`.
1. Hakuna kigezo cha kimataifa kinapaswa kutumika.

# --hints--

Unapaswa kuwa na kitendakazi cha `steamrollArray`.

```js
assert.isFunction(steamrollArray);
```

`steamrollArray([[["a"]], [["b"]]])` inapaswa kurudisha `["a", "b"]`.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])` inapaswa kurudisha `[1, 2, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])` inapaswa kurudisha `[1, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])` inapaswa kurudisha `[1, {}, 3, 4]`.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

Suluhisho lako halipaswi kutumia njia za `Array.prototype.flat()` au `Array.prototype.flatMap()`.

```js
const arr = [1,2,3,[1,2,3,[1,2,3]]];

const spyFlat = __helpers.spyOn(Array.prototype, 'flat');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlat.calls);
} finally {
  spyFlat.restore();
}

const spyFlatMap = __helpers.spyOn(Array.prototype, 'flatMap');
try {
  steamrollArray(arr);
  assert.isEmpty(spyFlatMap.calls);
} finally {
  spyFlatMap.restore();
}
```

Haupaswi kutumia vigezo vya kimataifa.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  const out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
