---
id: 64dc916138dbe03bdfe6cec6
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Мета гри "Вежа Ханоя" — перемістити всі диски на останній стрижень. Для цього потрібно дотримуватися трьох простих правил:

1. Можна рухати лише верхні диски
2. Можна рухати лише один диск за раз
3. Не можна класти більші диски на менші

Під вашим існуючим кодом оголосіть порожню функцію з іменем `move`. Пізніше ви будете використовувати цю функцію для переміщення дисків між стрижнями. Наразі, щоб уникнути помилок, використайте ключове слово `pass` всередині тіла функції.

# --hints--

Ви повинні оголосити порожню функцію з іменем `move`. Не забудьте використати ключове слово `pass` всередині тіла функції з правильним відступом.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
