---
id: 64dc8ea01436383a88256d7f
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Головоломка починається з того, що диски складені на першому стрижні за спаданням розміру, найменший диск зверху, а найбільший диск знизу. Вам потрібно заповнити свій перший список числами, які представляють різні розміри дисків.

Замість того, щоб додавати елементи вручну до першого списку, згенеруйте послідовність чисел, що зменшується від `3` до `1`, використовуючи функцію `range()`, і призначте її `rods['A']`. Тут `3` представляє найбільший диск на дні стопки, а `1` — найменший диск зверху стопки.

Синтаксис такий: `range(x, y, h)`, де `x` — початкове ціле число (включно), `y` — останнє ціле число (не включно), а `h` — різниця між числом і наступним у послідовності.

# --hints--

Ви повинні використати функцію `range()`, щоб призначити послідовність чисел `rods['A']`. Синтаксис виклику функції `range()` — `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Ви повинні використати функцію `range()`, щоб призначити послідовність чисел від `3` до `1` для `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
