---
id: 64dc9004f658183afa85f4f1
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер перевірте тип даних вашого ключа `'A'`, передавши його у функцію `type()` і виведіть результат у термінал.

# --hints--

Ви повинні передати свій ключ `'A'` у функцію `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Ви повинні вивести тип `rods['A']` за допомогою функції `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
