---
id: 64dc922df2919e3c38cead61
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

На початку вашого коду оголосіть змінну з іменем `NUMBER_OF_DISKS`, щоб зберігати число дисків, і надайте їй значення `3`. Потім замініть перший аргумент, переданий у функцію `range()`, на вашу нову змінну.

# --hints--

Ви повинні оголосити змінну з іменем `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

Змінна `NUMBER_OF_DISKS` має мати значення `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Ви повинні змінити значення першого аргументу функції `range()` на змінну `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
