---
id: 64dc90949ff85e3b37be40d0
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Функція `range()` повертає незмінну послідовність чисел. Як ви бачите, тип даних `rods['A']` — `range`, але ви хочете, щоб це був список.

Передайте свій виклик `range()` у функцію `list()`, щоб це зробити.

# --hints--

Ви повинні передати свій виклик `range()` у функцію `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
