---
id: 64dc8b1faf9c60304ca8a9b6
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви розв’яжете математичну головоломку, відому як вежа Ханоя. Головоломка складається з трьох стрижнів і кількох дисків різного діаметра.

Мета цієї головоломки — перемістити диски з першого стрижня на третій, дотримуючись певних правил, які забороняють класти більший диск на менший.

Почніть зі створення порожнього словника з назвою `rods`, щоб представити стрижні.

# --hints--

У вас має бути змінна з назвою `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

Змінна `rods` має бути порожнім словником.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
