---
id: 64dc97005bc7943e2978df0a
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер ви можете видалити свій виклик `print()`. Потім, всередині функції `move()`, видаліть ключове слово `pass` і виведіть вміст вашого словника `rods`.

# --hints--

У вашому коді не повинно бути `print(number_of_moves)`.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Ви повинні видалити ключове слово `pass` і вивести словник `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
