---
id: 64dc985e6720a23edac63b51
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Дозволені переміщення дисків між стрижнями мають повторюваний шаблон, який повторюється кожні три ходи. Наприклад, переміщення між стрижнями `A` і `C` дозволені на першому, четвертому та сьомому ході, коли остача від ділення номера ходу на 3 дорівнює 1.

Всередині раніше створеного циклу `for` замініть існуючий виклик `print()` на інструкцію `if`, яка виконується, коли `(i + 1) % 3 == 1`. У межах цієї інструкції `if` виведіть `f'Move {i + 1} allowed between {source} and {target}'` за допомогою f-рядка. Зверніть увагу, що `i + 1` — це номер ходу, оскільки `i` дорівнює нулю під час першої ітерації.

# --hints--

У вас має бути інструкція `if`, яка виконується, коли `(i + 1) % 3 == 1`.

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Виведіть рядок `f'Move {i + 1} allowed between {source} and {target}'` за допомогою f-рядка.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
