---
id: 64dc97ab25730b3e9de86ea5
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Наприкінці цього проєкту ви створите рекурсивне рішення головоломки "Вежа Ханоя", але зараз ви дослідите ітеративний підхід до цієї задачі.

Почніть із додавання `for` циклу у вашу функцію, який ітеруватиме по `number_of_moves` та виводитиме поточне число ітерації.

# --hints--

Вам слід написати `for` цикл, щоб ітерувати по кількості ходів. Для цього використовуйте функцію `range()`.

```js
({ test: () => {
    const rgs = [
        /for\s+\w+\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

Виводьте поточне число ходу на кожній ітерації.

```js
({ test: () => {
    const rgs = [
        /for\s+(\w+)\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:\s+print\s*\(\s*\1\s*\)/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
