---
id: 64dc92a9718fb53ca3d1bad7
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Головоломку вежа Ханоя можна розв’язати за 2<sup>n</sup> - 1 ходів, де n — кількість дисків. Оголосіть змінну з ім’ям `number_of_moves` і призначте цій змінній загальну кількість ходів.

Оператор степеня в Python — це `**`.

# --hints--

Ви повинні оголосити змінну з ім’ям `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Значення `number_of_moves` має бути виразом для обчислення кількості ходів.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
