---
id: 64ddd280d69fa818ca782d50
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Як ви бачите, диск 1 рухається туди й назад кожні три ходи. Це відбувається тому, що вам ще потрібно опрацювати переміщення між іншими стрижнями. Замість того, щоб повторювати той самий код, який ви писали в попередніх кроках, змінюючи стрижні, краще винести цей код у функцію, яку можна викликати в кожній умовній інструкції. Оголосіть порожню функцію з ім’ям `make_allowed_move()` і не забудьте про ключове слово `pass`.

# --hints--

Ви повинні оголосити порожню функцію з ім’ям `make_allowed_move`. Пам’ятайте використати ключове слово `pass` у тілі функції з правильним відступом.

```js
({ test: () => assert(runPython(`
      import inspect
      inspect.isfunction(make_allowed_move)
  `)) 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
