---
id: 64de1469fccce61940f498c4
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Словник `rods` має представляти три стрижні з їхніми дисками. Додайте до нього рядки `'A'`, `'B'` та `'C'` як ключі і встановіть кожен з них у порожній список.

# --hints--

Ваш словник `rods` повинен містити ключ `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Ваш словник `rods` повинен містити ключ `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Ваш словник `rods` повинен містити ключ `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` має бути порожнім списком.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
