---
id: 64df3f1011888113fbd3d81b
title: Крок 44
challengeType: 20
dashedName: step-44
---

# --description--

Щоб розв’язати головоломку за допомогою рекурсії, перш за все потрібно розбити початкову проблему на менші підпроблеми.

Фінальна конфігурація з `n` дисками, складеними на третій стрижень у спадному порядку, може бути отримана шляхом переміщення:

- `n - 1` дисків з початкового стрижня на допоміжний
- найбільшого диска з початкового на цільовий
- а потім `n - 1` дисків з допоміжного на цільовий.

Отже, перше, що має зробити `move` функція — це викликати саму себе з `n - 1` як першим аргументом. Але якщо спробувати зробити це без визначення базового випадку, ви отримаєте `RecursionError`. Це відбувається тому, що функція продовжує викликати себе безкінечно.

Перед вашим коментарем і вашим `print()` викликом додайте рекурсивний виклик функції з `n - 1` як першим аргументом і переконайтеся, що тіло функції виконується лише тоді, коли `n` більше за нуль. Поки що залиште інші аргументи в тому ж порядку.

# --hints--

Тіло `move` функції має починатися з `if` інструкції, яка спрацьовує, коли `n` більше за нуль.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Ви маєте викликати `move(n - 1, source, auxiliary, target)` всередині `if` інструкції. Не забудьте відступити ваш `print()` виклик.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
