---
id: 64df496c6a8ddf1b38db1ed6
title: Крок 50
challengeType: 20
dashedName: step-50
---

# --description--

У попередньому кроці ви написали код для переміщення найбільшого диска підзадачі на цільовий стрижень.

Тепер усе, що потрібно зробити — це додати ще один рекурсивний виклик, щоб перемістити `n - 1` дисків, які ви вже перемістили. Скопіюйте перший рекурсивний виклик і вставте його в кінець блоку `if`.

Зверніть увагу, що аргументи функції не в правильному порядку. Спробуйте визначити правильний порядок.

# --hints--

Ви повинні викликати `move(n - 1,  auxiliary, source, target)` у кінці вашої функції `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
