---
id: 64df47b32b92301a815d5ef8
title: Крок 49
challengeType: 20
dashedName: step-49
---

# --description--

Спочатку рекурсивний виклик, який ви щойно додали, працює з підзадачею переміщення `n - 1` дисків на другу стійку.

З цієї причини аргумент `target` відповідає вашій другій стійці, а аргумент `auxiliary` — третій стійці. Пам’ятайте, що вони будуть змінюватися місцями в процесі рекурсії.

Виправте порядок аргументів, помінявши місцями `target` і `auxiliary` у вашому рекурсивному виклику.

# --hints--

Вам слід змінити порядок аргументів у вашому виклику `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
