---
id: 650c6082e5586f9e3acfcd3b
title: Крок 52
challengeType: 20
dashedName: step-52
---

# --description--

Тепер видаліть словник `rods` і перетворіть його ключі на змінні `A`, `B` та `C` відповідно, зберігаючи їхні значення. Рефакторіть свій код, щоб відобразити ці зміни. Якщо ви бачите вихідні дані на терміналі, значить ви правильно виконали цей крок.

# --hints--

Ви повинні перетворити `rods['A']` на змінну з ім’ям `A`, зберігаючи її значення.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Ви повинні перетворити `rods['B']` на змінну з ім’ям `B`, зберігаючи її значення.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Ви повинні перетворити `rods['C']` на змінну з ім’ям `C`, зберігаючи її значення.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Ви повинні змінити свій виклик `print()`, щоб вивести `A`, `B`, `C`, замість об’єкта `rods`. Залиште символ нового рядка у виклику `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

У вашому коді має бути `target.append(source.pop())`.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
