---
id: 68da6138d70960c197ec5ef0
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Jetzt implementieren Sie die Haupt-BFS-Schleife. Erstellen Sie eine `while`-Schleife, die so lange läuft, wie die `queue` nicht leer ist (während `queue` zu `True` ausgewertet wird).

Geben Sie innerhalb der Schleife `queue` aus.

# --hints--

Sie sollten eine `while`-Schleife in Ihrer `gen_parentheses`-Funktion haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

Ihre `while`-Schleife sollte die Bedingung `queue` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

Sie sollten `queue` innerhalb Ihrer `while`-Schleife ausgeben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
