---
id: 68da6138d70960c197ec5ef1
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Jetzt müssen Sie erkennen, wann Sie eine vollständige Kombination von runden Klammern erstellt haben. Eine vollständige Kombination hat eine Länge, die dem Doppelten der Anzahl der Paare entspricht (da jedes Paar eine öffnende und eine schließende Klammer beiträgt).

Fügen Sie innerhalb Ihrer `while`-Schleife eine `if`-Anweisung hinzu, um zu prüfen, ob `len(current) == 2 * pairs` zutrifft. Wenn diese Bedingung wahr ist, fügen Sie `current` der `result`-Liste hinzu.

# --hints--

Sie sollten eine `if`-Anweisung innerhalb Ihrer `while`-Schleife haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

Ihre `if`-Anweisung sollte prüfen, ob `len(current) == 2 * pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

Sie sollten `current` der `result`-Liste innerhalb Ihrer neuen `if`-Anweisung hinzufügen.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
--fcc-editable-region--
        
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
