---
id: 68da6138d70960c197ec5eee
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Jetzt richten Sie die Datenstruktur ein, um Ihre Ergebnisse zu speichern. Erstellen Sie eine Variable namens `result` und initialisieren Sie sie als leere Liste. Diese Liste wird alle gültigen Kombinationen von runden Klammern speichern, die Sie erzeugen.

Aktualisieren Sie Ihre return-Anweisung, sodass `result` statt einer leeren Liste zurückgegeben wird.

# --hints--

Sie sollten eine Variable namens `result` in Ihrer `gen_parentheses`-Funktion deklarieren.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

Sie sollten `result` als leere Liste initialisieren.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

Sie sollten am Ende Ihrer Funktion `result` zurückgeben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
    return []
--fcc-editable-region--
```
