---
id: 68da6138d70960c197ec5eed
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Als Nächstes müssen Sie überprüfen, ob die Anzahl der Paare mindestens eins beträgt, da Sie keine Kombinationen von runden Klammern mit null oder negativen Paaren erzeugen können.

Fügen Sie eine weitere `if`-Anweisung hinzu, um zu prüfen, ob `pairs` kleiner als `1` ist. Wenn diese Bedingung zutrifft, geben Sie den String `The number of pairs should be at least 1` zurück.

# --hints--

Sie sollten eine zweite `if`-Anweisung in Ihrer `gen_parentheses`-Funktion haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

Ihre zweite `if`-Anweisung sollte prüfen, ob `pairs` kleiner als `1` ist.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

Sie sollten den String `The number of pairs should be at least 1` zurückgeben, wenn `pairs` kleiner als `1` ist.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
