---
id: 68da6138d70960c197ec5eec
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Bevor Sie den Kernalgorithmus implementieren, müssen Sie die Eingabe validieren. Der `pairs`-Parameter sollte eine ganze Zahl sein, da er die Anzahl der zu erzeugenden Klammerpaare darstellt.

Fügen Sie am Anfang Ihrer Funktion eine `if`-Anweisung hinzu, um zu prüfen, ob `pairs` keine ganze Zahl ist. Verwenden Sie dafür die `isinstance()`-Funktion.

Wenn die Bedingung zutrifft, geben Sie den String `The number of pairs should be an integer` zurück.

# --hints--

Sie sollten eine `if`-Anweisung in Ihrer `gen_parentheses`-Funktion haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

Ihre `if`-Anweisung sollte prüfen, ob `pairs` keine ganze Zahl ist, und dabei `not isinstance(pairs, int)` verwenden.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

Sie sollten den String `The number of pairs should be an integer` zurückgeben, wenn die Eingabe keine ganze Zahl ist.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
--fcc-editable-region--
    
--fcc-editable-region--
    
    return []
```
