---
id: 68da6138d70960c197ec5eef
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Für die Breitensuche verwenden Sie eine Warteschlange, um verschiedene Zustände zu verfolgen, während Sie die Kombinationen der runden Klammern aufbauen. Jeder Zustand wird als Tuple mit drei Elementen dargestellt:

- Der aktuell aufgebaute String
- Die Anzahl der bisher verwendeten öffnenden runden Klammern
- Die Anzahl der bisher verwendeten schließenden runden Klammern

Erstellen Sie eine Variable namens `queue` und initialisieren Sie sie mit einer Liste, die ein Tuple enthält: `('', 0, 0)`. Dies stellt den Anfangszustand mit einem leeren String und null verwendeten Klammern dar.

# --hints--

Sie sollten eine Variable namens `queue` in Ihrer `gen_parentheses`-Funktion deklarieren.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

Sie sollten `queue` mit einer Liste initialisieren, die das Tuple `('', 0, 0)` enthält.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
--fcc-editable-region--
    
--fcc-editable-region--
    result = []
    return result
```
