---
id: 68da604b8b6012a33e2344fb
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop implementieren Sie eine Funktion, die alle gültigen Kombinationen von runden Klammern mithilfe eines Breitensuche-Ansatzes generiert. Zum Beispiel sind die gültigen Kombinationen von zwei Paaren runder Klammern `(())` und `()()`.

Beginnen Sie damit, eine Funktion mit dem Namen `gen_parentheses` zu erstellen, die einen einzelnen Parameter `pairs` hat. Geben Sie vorerst eine leere Liste aus der Funktion zurück.

# --hints--

Sie sollten eine Funktion mit dem Namen `gen_parentheses` definieren.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

Ihre `gen_parentheses`-Funktion sollte einen einzelnen Parameter namens `pairs` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

Ihre Funktion sollte eine leere Liste zurückgeben.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
