---
id: 68da6138d70960c197ec5ef3
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Jetzt müssen Sie das Hinzufügen von schließenden runden Klammern behandeln. Die wichtigste Regel ist, dass Sie eine schließende Klammer nur hinzufügen dürfen, wenn dadurch das Gleichgewicht erhalten bleibt, das heißt, es müssen mehr öffnende runde Klammern verwendet worden sein als schließende.

Fügen Sie innerhalb des `else`-Blocks eine zweite `if`-Anweisung hinzu, um zu prüfen, ob `closes_used < opens_used`.

Wenn diese Bedingung wahr ist, fügen Sie dem `queue` ein weiteres neues Tuple hinzu: `(current + ')', opens_used, closes_used + 1)`. Dies stellt den Zustand nach dem Hinzufügen einer schließenden Klammer dar.

# --hints--

Sie sollten eine zweite `if`-Anweisung innerhalb Ihres `else`-Blocks haben.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

Ihre zweite verschachtelte `if`-Anweisung sollte prüfen, ob `closes_used < opens_used`.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

Sie sollten `(current + ')', opens_used, closes_used + 1)` zur Warteschlange hinzufügen, wenn die Bedingung wahr ist.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
--fcc-editable-region--
            
--fcc-editable-region--

    return result

print(gen_parentheses(1))
```
