---
id: 68da6138d70960c197ec5ef2
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wenn der aktuelle String noch nicht vollständig ist, müssen Sie die nächsten möglichen Zustände erkunden. Fügen Sie eine `else`-Klausel zu Ihrer `if`-Anweisung hinzu.

Innerhalb des `else`-Blocks behandeln Sie das Hinzufügen von öffnenden runden Klammern. Fügen Sie eine `if`-Anweisung hinzu, um zu prüfen, ob `opens_used < pairs` ist. So stellen Sie sicher, dass Sie nicht mehr öffnende runde Klammern verwenden, als erlaubt sind.

Wenn diese Bedingung wahr ist, fügen Sie ein neues Tuple an `queue` an: `(current + '(', opens_used + 1, closes_used)`. Dies stellt den Zustand nach dem Hinzufügen einer öffnenden runden Klammer dar.

# --hints--

Sie sollten eine `else`-Klausel für Ihre `if`-Anweisung innerhalb der `while`-Schleife haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

Sie sollten eine `if`-Anweisung innerhalb des `else`-Blocks haben, die prüft, ob `opens_used < pairs`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

Sie sollten `(current + '(', opens_used + 1, closes_used)` zur Warteschlange hinzufügen, wenn die Bedingung wahr ist.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        
--fcc-editable-region--
    
    return result

print(gen_parentheses(1))
```
