---
id: 68e00c5349cbddfef6dbc053
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Innerhalb Ihrer `while`-Schleife verwenden Sie `queue.pop(0)`, um das erste Element aus der Warteschlange zu entfernen und zu erhalten. Dies implementiert das First-in-first-out-(FIFO-)Verhalten, das für BFS charakteristisch ist.

Packen Sie dieses Tuple in drei Variablen aus: `current`, `opens_used` und `closes_used`.

# --hints--

Sie sollten `queue.pop(0)` verwenden, um das erste Element aus der Warteschlange zu erhalten.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

Sie sollten das Ergebnis in drei Variablen auspacken: `current`, `opens_used` und `closes_used`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    while queue:
        print(queue)
--fcc-editable-region--
        
--fcc-editable-region--

    return result
```
