---
id: 673b02b03134b04637bf7055
title: Construir um Rodapé Reutilizável
challengeType: 25
dashedName: build-a-reusable-footer
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve exportar um componente `Footer`.
2. Seu componente `Footer` deve retornar um elemento `footer` que contenha todos os outros elementos.
3. Your `footer` element should have:

   - Pelo menos três listas não ordenadas, cada uma com pelo menos dois itens de lista.
   - Pelo menos um elemento de parágrafo com o símbolo de copyright `©`.
   - Pelo menos três links com o valor `href` definido como `#` e o conteúdo do link definido como um ícone ou texto de sua escolha.

# --hints--

Você deve exportar um componente `Footer`.

```js
assert.isFunction(window.index.Footer);
```

Seu componente `Footer` deve retornar um elemento `footer`.

```js
const footer = document.getElementById('root').firstElementChild;
assert.equal(footer.tagName, 'FOOTER');
```

Seu componente `Footer` deve conter apenas um único elemento `footer`.

```js
assert.lengthOf(document.querySelectorAll('footer'), 1);
```

Seu elemento `footer` não deve ter irmãos.

```js
assert.lengthOf(document.getElementById('root').children, 1);
```

Seu elemento `footer` deve conter pelo menos três listas não ordenadas.

```js
assert.isAtLeast(document.querySelectorAll('ul').length, 3);
```

Cada uma das suas listas não ordenadas deve conter pelo menos dois itens de lista.

```js
const uls = [...document.querySelectorAll('ul')];
uls.forEach(ul => {
    assert.isAtLeast(ul.querySelectorAll('li').length, 2);
  }
);
```

Seu `footer` deve conter pelo menos um elemento de parágrafo.

```js
assert.isAtLeast(document.querySelectorAll('p').length, 1);
```

Você deve ter um símbolo de copyright (`©`) dentro de um dos seus parágrafos.

```js
const texts = [...document.querySelectorAll('p')].map(p => p.textContent);
const hasCopyright = texts.some(text => text.includes('©'));
assert.isTrue(hasCopyright);
```

Seu `footer` deve conter pelo menos três links com o valor do `href` definido como `#`.

```js
const links = document.querySelectorAll('a[href="#"]');
assert.isAtLeast(links.length, 3);
```

Nenhum dos seus links deve estar vazio.

```js
const links = [...document.querySelectorAll('a[href="#"]')];
links.forEach(link => {
    assert.isNotEmpty(link.textContent);
  }
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css

```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Reusable Footer Component</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { Footer } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<Footer />);
    </script>
  </body>
</html>
```

```css
.footer {
  background-color: #e3e1eb;
  padding: 20px;
  text-align: center;
  border-top: 1px solid #e5e7eb;
  font-family: Arial, sans-serif;
}

.footer-section {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  justify-content: center;
  margin-bottom: 20px;
  align-items: center;
}

.footer ul {
  list-style: none;
  padding: 0;
  margin: 0;
}

.footer li {
  margin-bottom: 10px;
}

.footer a {
  text-decoration: none;
  color: #1f2937;
}

.footer a:hover {
  color: #6366f1;
}

.footer p {
  color: #6b7280;
  margin-top: 20px;
}

.footer-icons {
  margin-top: 20px;
  display: flex;
  justify-content: center;
}

.footer-icons a {
  margin: 0 10px;
  font-size: 24px;
  color: #6b7280;
  text-decoration: none;
}

.footer-icons a:hover {
  color: #6366f1;
}
```

```jsx
export const Footer = () => (
  <footer className='footer'>
    <div className='footer-section'>
      <ul>
        <li>
          <a href='#'>Fitness Dashboard</a>
        </li>
        <li>
          <a href='#'>Services</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Watch Videos</a>
        </li>
        <li>
          <a href='#'>Discord</a>
        </li>
      </ul>
      <ul>
        <li>
          <a href='#'>Privacy Policy</a>
        </li>
        <li>
          <a href='#'>Terms & Conditions</a>
        </li>
      </ul>
    </div>
    <p>© 2024 Fitness Dashboard. All Rights Reserved.</p>
    <div className='footer-icons'>
      <a href='#'>&#x1F3AE;</a>
      <a href='#'>&#x1F426;</a>
      <a href='#'>&#x1F4BB;</a>
      <a href='#'>&#x1F3C0;</a>
    </div>
  </footer>
);
```
