---
id: 587d7fb7367417b2b2512c0d
title: usa model.findById() per cercare nel tuo database tramite _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Quando salvi un documento, MongoDB aggiunge automaticamente il campo `_id` e lo imposta su una chiave alfanumerica unica. La ricerca tramite `_id` è un’operazione estremamente frequente, quindi Mongoose fornisce un metodo dedicato per questo.

# --instructions--

Modifica `findPersonById` per trovare la persona che ha `_id` specificato, usando `Model.findById() -> Person`. Usa l’argomento della funzione `personId` come chiave di ricerca.

# --hints--

Trovare un elemento tramite Id dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
