---
id: 587d7fb7367417b2b2512c0b
title: Usa model.find() per cercare nel tuo database
challengeType: 2
forumTopicId: 301543
dashedName: use-model-find-to-search-your-database
---

# --description--

Nel suo utilizzo più semplice, `Model.find()` accetta come primo argomento un documento di query (un oggetto JSON), poi una callback. Restituisce un array di corrispondenze. Supporta un'ampia gamma di opzioni di ricerca. Leggi di più nella documentazione.

# --instructions--

Modifica la funzione `findPeopleByName` per trovare tutte le persone che hanno un nome specifico, usando <code>Model.find() -\> [Person]</code>

Usa l'argomento della funzione `personName` come chiave di ricerca.

# --hints--

Trovare tutti gli elementi corrispondenti a un criterio dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/find-all-by-name', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'r@nd0mN4m3',
      age: 24,
      favoriteFoods: ['pizza']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data[0].name,
    'r@nd0mN4m3',
    'item.name is not what expected'
  );
  assert.equal(data[0].__v, 0, 'The item should be not previously edited');
```
