---
id: 587d7fb8367417b2b2512c11
title: eliminare molti documenti con model.remove()
challengeType: 2
forumTopicId: 301538
dashedName: delete-many-documents-with-model-remove
---

# --description--

`Model.remove()` è utile per eliminare tutti i documenti che corrispondono ai criteri specificati.

# --instructions--

Modifica la funzione `removeManyPeople` per eliminare tutte le persone il cui nome è contenuto nella variabile `nameToRemove`, usando `Model.remove()`. Passala a un documento query con il campo `name` impostato, e una callback.

**Nota:** Il `Model.remove()` non restituisce il documento eliminato, ma un oggetto JSON contenente l'esito dell'operazione e il numero di elementi interessati. Non dimenticare di passarla alla callback `done()`, perché la usiamo nei test.

# --hints--

Eliminare molti elementi contemporaneamente dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/remove-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Mary', age: 16, favoriteFoods: ['lollipop'] },
      { name: 'Mary', age: 21, favoriteFoods: ['steak'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(!!data.ok, 'The mongo stats are not what expected');
  assert.equal(
    data.n,
    2,
    'The number of items affected is not what expected'
  );
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
