---
id: 587d7fb8367417b2b2512c0e
title: 'Eseguire aggiornamenti classici trovando, modificando e poi salvando'
challengeType: 2
forumTopicId: 301541
dashedName: perform-classic-updates-by-running-find-edit-then-save
---

# --description--

Ai tempi andava così: se volevi modificare un documento e poterlo usare in qualche modo (ad esempio, rimandandolo in una risposta del server), dovevi fare così. Mongoose ha un metodo dedicato per aggiornare: `Model.update()`. È collegato al driver mongo di basso livello. Può modificare in blocco molti documenti che corrispondono a certi criteri, ma non restituisce il documento aggiornato, solo un messaggio di 'stato'. Inoltre, rende difficile la validazione del modello, perché chiama direttamente il driver mongo.

# --instructions--

Modifica la funzione `findEditThenSave` per trovare una persona tramite `_id` (usa uno qualsiasi dei metodi sopra) con il parametro `personId` come chiave di ricerca. Aggiungi `"hamburger"` alla lista di `favoriteFoods` della persona (puoi usare `Array.push()`). Poi — all’interno della callback di find — `save()` il `Person` aggiornato.

**Nota:** Questo può essere complicato se nel tuo Schema hai dichiarato `favoriteFoods` come Array, senza specificarne il tipo (cioè `[String]`). In quel caso, `favoriteFoods` è di tipo Mixed per impostazione predefinita, e devi segnalarlo manualmente come modificato usando `document.markModified('edited-field')`. Consulta il nostro <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">articolo su Mongoose</a>.

# --hints--

Trovare, modificare e aggiornare un elemento dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/find-edit-save', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Poldo',
      age: 40,
      favoriteFoods: ['spaghetti']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Poldo', 'item.name is not what is expected');
  assert.equal(data.age, 40, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['spaghetti', 'hamburger'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 1, 'The item should be previously edited');
```
