---
id: 587d7fb7367417b2b2512c0a
title: Crea molti record con model.create()
challengeType: 2
forumTopicId: 301537
dashedName: create-many-records-with-model-create
---

# --description--

A volte serve creare molte istanze dei tuoi modelli, per esempio quando si popola un database con dati iniziali. `Model.create()` prende come primo argomento un array di oggetti come `[{name: 'John', ...}, {...}, ...]` e li salva tutti nel database.

# --instructions--

Modifica la funzione `createManyPeople` per creare molte persone usando `Model.create()` con l'argomento `arrayOfPeople`.

**Nota:** Puoi riutilizzare il modello che hai istanziato nell'esercizio precedente.

# --hints--

Creare molti elementi nel database contemporaneamente dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/create-many-people', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'John', age: 24, favoriteFoods: ['pizza', 'salad'] },
      { name: 'Mary', age: 21, favoriteFoods: ['onions', 'chicken'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an array');
  assert.equal(
    data.length,
    2,
    'the response does not contain the expected number of items'
  );
  assert.equal(data[0].name, 'John', 'The first item is not correct');
  assert.equal(
    data[0].__v,
    0,
    'The first item should be not previously edited'
  );
  assert.equal(data[1].name, 'Mary', 'The second item is not correct');
  assert.equal(
    data[1].__v,
    0,
    'The second item should be not previously edited'
  );
```
