---
id: 587d7fb6367417b2b2512c06
title: installare e configurare mongoose
challengeType: 2
forumTopicId: 301540
dashedName: install-and-set-up-mongoose
---

# --description--

Per affrontare queste sfide dovrai scrivere il tuo codice usando uno dei seguenti metodi:

- Clona <a href="https://github.com/freeCodeCamp/boilerplate-mongomongoose/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa le sfide in locale.
- Usa un site builder a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

In questa sfida, configurerai un database MongoDB Atlas e importerai i pacchetti necessari per connetterti a esso.

Segui <a href='https://www.freecodecamp.org/news/get-started-with-mongodb-atlas/' target="_blank" rel="noopener noreferrer nofollow">questa guida</a> per configurare un database ospitato su MongoDB Atlas.

# --instructions--

`mongoose@^5.11.15` è stato aggiunto al file `package.json` del tuo progetto. Per prima cosa, importa mongoose come `mongoose` in `myApp.js`. Poi, crea un file `.env` e aggiungi una variabile `MONGO_URI`. Il suo valore deve essere l’URI del tuo database MongoDB Atlas. Assicurati di racchiudere l’URI tra virgolette singole o doppie e ricorda che non puoi usare spazi intorno al `=` nelle variabili d’ambiente. Per esempio, `MONGO_URI='VALUE'`.

Quando hai finito, connettiti al database chiamando il metodo `connect` all’interno del file `myApp.js` usando la seguente sintassi:

```js
mongoose.connect(<Your URI>, { useNewUrlParser: true, useUnifiedTopology: true });
```

# --hints--

La dipendenza "mongoose version ^5.11.15" dovrebbe essere presente in package.json

```js
  const response = await fetch(code + '/_api/file/package.json');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.text();
  var packJson = JSON.parse(data);
  assert.property(packJson.dependencies, 'mongoose');
  assert.match(
    packJson.dependencies.mongoose,
    /^\^5\.11\.15/,
    'Wrong version of "mongoose". It should be ^5.11.15'
  );
```

"mongoose" dovrebbe essere connesso a un database

```js
  const response = await fetch(code + '/_api/is-mongoose-ok');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isTrue(data.isMongooseOk, 'mongoose is not connected');
```
