---
id: 587d7fb9367417b2b2512c12
title: concatenare i helper di query di ricerca per restringere i risultati
challengeType: 2
forumTopicId: 301533
dashedName: chain-search-query-helpers-to-narrow-search-results
---

# --description--

Se non passi la callback come ultimo argomento a `Model.find()` (o agli altri metodi di ricerca), la query non viene eseguita. Puoi archiviare la query in una variabile per usarla in seguito. Questo tipo di oggetto ti permette di costruire una query usando la sintassi di concatenamento. La ricerca effettiva nel db viene eseguita quando concaten finalmente il metodo `.exec()`. Devi sempre passare la tua callback a quest’ultimo metodo. Esistono molti helper per le query, qui useremo quelli più comuni.

# --instructions--

Modifica la funzione `queryChain` per trovare le persone a cui piace il cibo specificato dalla variabile chiamata `foodToSearch`. Ordinale per `name`, limita i risultati a due documenti e nascondi la loro età. Concatenare `.find()`, `.sort()`, `.limit()`, `.select()` e poi `.exec()`. Passa la callback `done(err, data)` a `exec()`.

# --hints--

Il concatenamento degli helper di query dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/query-tools', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify([
      { name: 'Pablo', age: 26, favoriteFoods: ['burrito', 'hot-dog'] },
      { name: 'Bob', age: 23, favoriteFoods: ['pizza', 'nachos'] },
      { name: 'Ashley', age: 32, favoriteFoods: ['steak', 'burrito'] },
      { name: 'Mario', age: 51, favoriteFoods: ['burrito', 'prosciutto'] }
    ])
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isArray(data, 'the response should be an Array');
  assert.equal(
    data.length,
    2,
    'the data array length is not what expected'
  );
  assert.notProperty(
    data[0],
    'age',
    'The returned first item has too many properties'
  );
  assert.equal(
    data[0].name,
    'Ashley',
    'The returned first item name is not what expected'
  );
  assert.notProperty(
    data[1],
    'age',
    'The returned second item has too many properties'
  );
  assert.equal(
    data[1].name,
    'Mario',
    'The returned second item name is not what expected'
  );
```
