---
id: 587d7fb7367417b2b2512c0c
title: usa model.findOne() per restituire un singolo documento corrispondente dal tuo database
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` si comporta come `Model.find()`, ma restituisce solo un documento (non un array), anche se ci sono più elementi. È particolarmente utile quando si fa una ricerca per proprietà che hai dichiarato come uniche.

# --instructions--

Modifica la funzione `findOneByFood` per trovare solo una persona che abbia un certo cibo tra i preferiti, usando `Model.findOne() -> Person`. Usa l'argomento della funzione `food` come chiave di ricerca.

# --hints--

Trovare un elemento dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```
