---
id: 587d7fb6367417b2b2512c09
title: Crea e salva un record di un modello
challengeType: 2
forumTopicId: 301536
dashedName: create-and-save-a-record-of-a-model
---

# --description--

In questa sfida dovrai creare e salvare un record di un modello.

# --instructions--

All'interno della funzione `createAndSavePerson`, crea un'istanza del documento usando il costruttore del modello `Person` che hai costruito prima. Passa al costruttore un oggetto con i campi `name`, `age` e `favoriteFoods`. I loro tipi devono corrispondere a quelli nel `personSchema`. Poi, chiama il metodo `document.save()` sull'istanza del documento restituita. Passa a questo metodo una callback usando la convenzione Node. Questo è un pattern comune; tutti i metodi CRUD seguenti prendono una funzione di callback come ultimo argomento.

```js
/* Example */

// ...
person.save(function(err, data) {
  //   ...do your stuff here...
});
```

# --hints--

Creare e salvare un elemento nel database dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/create-and-save-person');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.isString(data.name, '"item.name" should be a String');
  assert.isNumber(data.age, '28', '"item.age" should be a Number');
  assert.isArray(
    data.favoriteFoods,
    '"item.favoriteFoods" should be an Array'
  );
  assert.equal(data.__v, 0, 'The db item should be not previously edited');
```
