---
id: 587d7fb8367417b2b2512c10
title: Eliminare un documento usando model.findByIdAndRemove
challengeType: 2
forumTopicId: 301539
dashedName: delete-one-document-using-model-findbyidandremove
---

# --description--

`findByIdAndRemove` e `findOneAndRemove` sono simili ai precedenti metodi di aggiornamento. Passano il documento eliminato al database. Come sempre, usa l'argomento della funzione `personId` come chiave di ricerca.

# --instructions--

Modifica la funzione `removeById` per eliminare una persona tramite il `_id` della persona. Dovresti usare uno dei metodi `findByIdAndRemove()` o `findOneAndRemove()`.

# --hints--

Eliminare un elemento dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/remove-one-person', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Jason Bourne',
      age: 36,
      favoriteFoods: ['apples']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Jason Bourne', 'item.name is not what expected');
  assert.equal(data.age, 36, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['apples'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0);
  assert.equal(data.count, 0, 'the db items count is not what expected');
```
