---
id: 587d7fb6367417b2b2512c07
title: Crea un modello
challengeType: 2
forumTopicId: 301535
dashedName: create-a-model
---

# --description--

**C**RUD parte I - CREARE

Prima di tutto, serve uno schema. Ogni schema corrisponde a una collezione MongoDB. Definisce la forma dei documenti all’interno di quella collezione. Gli schemi sono i mattoni per i modelli. Possono essere annidati per creare modelli complessi, ma in questo caso manterremo le cose semplici. Un modello ti permette di creare istanze dei tuoi oggetti, chiamate documenti.

Nei server, le interazioni con il database avvengono nelle funzioni gestore. Queste funzioni vengono eseguite quando succede un evento (ad esempio qualcuno colpisce un endpoint della tua API). Seguiremo lo stesso approccio in questi esercizi. La funzione `done()` è una callback che ci dice che possiamo procedere dopo aver completato un’operazione asincrona come inserire, cercare, aggiornare o eliminare. Segue la convenzione Node, e dovrebbe essere chiamata come `done(null, data)` in caso di successo, o `done(err)` in caso di errore.

Attenzione - Quando si interagisce con servizi remoti, possono verificarsi errori!

```js
/* Example */

const someFunc = function(done) {
  //... do something (risky) ...
  if (error) return done(error);
  done(null, result);
};
```

# --instructions--

Crea uno schema persona chiamato `personSchema` con la seguente forma:

* Un campo `name` obbligatorio di tipo `String`
* Un campo `age` di tipo `Number`
* Un campo `favoriteFoods` di tipo `[String]`

Usa i tipi di schema base di Mongoose. Se vuoi puoi anche aggiungere altri campi, usare validatori semplici come required o unique, e impostare valori predefiniti. Consulta il nostro <a href="https://www.freecodecamp.org/news/introduction-to-mongoose-for-mongodb-d2a7aa593c57/" target="_blank" rel="noopener noreferrer nofollow">articolo su Mongoose</a>.

Ora, crea un modello dallo `personSchema` e assegnalo alla variabile esistente `Person`.

# --hints--

Creare un’istanza da uno schema mongoose dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/mongoose-model', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Mike',
      age: 28,
      favoriteFoods: ['pizza', 'cheese']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Mike', '"model.name" is not what expected');
  assert.equal(data.age, '28', '"model.age" is not what expected');
  assert.isArray(
    data.favoriteFoods,
    '"model.favoriteFoods" is not an Array'
  );
  assert.include(
    data.favoriteFoods,
    'pizza',
    '"model.favoriteFoods" does not include the expected items'
  );
  assert.include(
    data.favoriteFoods,
    'cheese',
    '"model.favoriteFoods" does not include the expected items'
  );
```
