---
id: 587d7fb8367417b2b2512c0f
title: Eseguire nuovi aggiornamenti su un documento usando model.findOneAndUpdate()
challengeType: 2
forumTopicId: 301542
dashedName: perform-new-updates-on-a-document-using-model-findoneandupdate
---

# --description--

Le versioni recenti di Mongoose hanno metodi per semplificare l’aggiornamento dei documenti. Alcune funzionalità più avanzate (ad esempio i pre/post hook, la validazione) si comportano in modo diverso con questo approccio, quindi il metodo classico è ancora utile in molte situazioni. `findByIdAndUpdate()` può essere usato per la ricerca tramite id.

# --instructions--

Modifica la funzione `findAndUpdate` per trovare una persona tramite `Name` e impostare l’età della persona su `20`. Usa il parametro della funzione `personName` come chiave di ricerca.

**Nota:** Devi restituire il documento aggiornato. Per farlo, devi passare il documento delle opzioni `{ new: true }` come terzo argomento a `findOneAndUpdate()`. Per impostazione predefinita, questi metodi restituiscono l’oggetto non modificato.

# --hints--

findOneAndUpdate un elemento dovrebbe riuscire

```js
  const response = await fetch(code + '/_api/find-one-update', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Dorian Gray',
      age: 35,
      favoriteFoods: ['unknown']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Dorian Gray', 'item.name is not what expected');
  assert.equal(data.age, 20, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['unknown'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(
    data.__v,
    0,
    'findOneAndUpdate does not increment version by design!'
  );
```
