---
id: 587d7dbd367417b2b2512bb6
title: Unda CSS inayoweza kutumika tena kwa kutumia mixins
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

Katika Sass, <dfn>mixin</dfn> ni kundi la tamko za CSS zinazoweza kutumika tena katika karatasi ya mtindo mzima. Ufafanuzi huanza na kanuni ya `@mixin` ya at, ikifuatiwa na jina la kawaida. Unatumia mixin kwa kutumia kanuni ya `@include` ya at.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

Inatengeneza:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

Mixins zako pia zinaweza kuchukua hoja, ambayo inaruhusu tabia zao kubinafsishwa. Hoja hizo zinahitajika wakati wa kutumia mixin.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

Unaweza kufanya hoja kuwa za hiari kwa kuwapa vigezo thamani za msingi.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

Andika mixin iitwayo `shape` na ipe vigezo 3: `$w`, `$h`, na `$bg-color`.

Tumia mixin ya `shape` kumpa kipengele cha `#square` upana na urefu wa `50px`, na rangi ya nyuma `red`. Kwa kipengele cha `#rect-a` ongeza upana wa `100px`, urefu wa `50px`, na rangi ya nyuma `blue`. Mwishowe, kwa kipengele cha `#rect-b` ongeza upana wa `50px`, urefu wa `100px`, na rangi ya nyuma `orange`.

# --hints--

Unapaswa kutangaza mixin iitwayo `shape` yenye vigezo 3: `$w`, `$h`, na `$bg-color`.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

Mixin yako inapaswa kujumuisha sifa ya `width` inayotumia kigezo cha `$w`.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

Mixin yako inapaswa kujumuisha sifa ya `height` inayotumia kigezo cha `$h`.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

Mixin yako inapaswa kujumuisha sifa ya `background-color` inayotumia kigezo cha `$bg-color`.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

Unapaswa kubadilisha mitindo ndani ya kichaguzi cha `#square` kwa wito wa mixin ya `shape` ukitumia neno kuu la `@include`. Weka upana na urefu wa `50px`, na rangi ya nyuma `red`.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

Unapaswa kubadilisha mitindo ndani ya kichaguzi cha `#rect-a` kwa wito wa mixin ya `shape` ukitumia neno kuu la `@include`. Weka upana wa `100px`, urefu wa `50px`, na rangi ya nyuma `blue`.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

Unapaswa kubadilisha mitindo ndani ya kichaguzi cha `#rect-b` kwa wito wa mixin ya `shape` ukitumia neno kuu la `@include`. Weka upana wa `50px`, urefu wa `100px`, na rangi ya nyuma `orange`.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
