---
id: 587d7dbf367417b2b2512bbb
title: Tumia mtindo hadi hali itakapokamilika kwa @while
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

Amri ya `@while` ni chaguo lenye utendaji sawa na mzunguko wa JavaScript `while`. Inaunda sheria za CSS hadi hali fulani itakapokamilika.

Changamoto ya `@for` ilitoa mfano wa kuunda mfumo rahisi wa gridi. Hii pia inaweza kufanya kazi na `@while`.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

Kwanza, fafanua kigezo `$x` na uweke thamani yake kuwa 1. Kisha, tumia amri ya `@while` kuunda mfumo wa gridi *wakati* `$x` ni chini ya 13. Baada ya kuweka sheria ya CSS kwa `width`, `$x` inaongezwa kwa 1 ili kuepuka mzunguko usioisha.

# --instructions--

Tumia `@while` kuunda mfululizo wa madarasa ya vitu yenye `font-sizes` tofauti.

Kuwa na madarasa 5 tofauti kutoka `text-1` hadi `text-5`. Kisha weka `font-size` kuwa `15px` mara nambari ya kielezo cha sasa. Hakikisha kuepuka mzunguko usioisha!

# --hints--

Msimbo wako unapaswa kutumia amri ya `@while`.

```js
assert(code.match(/@while /g));
```

Msimbo wako unapaswa kutumia kigezo cha kielezo kinachoanza na thamani ya 1.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

Msimbo wako unapaswa kuongeza thamani ya kigezo cha kuhesabu.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

Darasa lako la `.text-1` linapaswa kuwa na `font-size` ya `15px`.

```js
assert($('.text-1').css('font-size') == '15px');
```

Darasa lako la `.text-2` linapaswa kuwa na `font-size` ya `30px`.

```js
assert($('.text-2').css('font-size') == '30px');
```

Darasa lako la `.text-3` linapaswa kuwa na `font-size` ya `45px`.

```js
assert($('.text-3').css('font-size') == '45px');
```

Darasa lako la `.text-4` linapaswa kuwa na `font-size` ya `60px`.

```js
assert($('.text-4').css('font-size') == '60px');
```

Darasa lako la `.text-5` linapaswa kuwa na `font-size` ya `75px`.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
