---
id: 587d7dbf367417b2b2512bba
title: Tumia @each kupitia vitu katika orodha
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

Changamoto ya mwisho ilionyesha jinsi agizo la `@for` linavyotumia thamani ya kuanzia na kumalizia kuzungusha mara fulani. Sass pia hutoa agizo la `@each` ambalo huzungusha kila kipengee katika orodha au ramani. Kila mzunguko, kigezo hupatiwa thamani ya sasa kutoka orodha au ramani.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

Ramani ina sintaksia tofauti kidogo. Hapa kuna mfano:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

Kumbuka kuwa kigezo cha `$key` kinahitajika kwa ajili ya kurejea funguo katika ramani. Vinginevyo, CSS iliyokusanywa ingekuwa na `color1`, `color2`... ndani yake. Mifano yote miwili ya msimbo hapo juu hubadilishwa kuwa CSS ifuatayo:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

Andika agizo la `@each` linalopita kwenye orodha: `blue, black, red` na kuweka kila kigezo kuwa darasa la `.color-bg`, ambapo sehemu ya `color` hubadilika kwa kila kipengee kuwa rangi husika. Kila darasa linapaswa kuweka `background-color` kwa rangi husika pia.

# --hints--

Msimbo wako unapaswa kutumia agizo la `@each`.

```js
assert(code.match(/@each /g));
```

Darasa lako la `.blue-bg` linapaswa kuwa na `background-color` ya bluu.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

Darasa lako la `.black-bg` linapaswa kuwa na `background-color` ya nyeusi.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

Darasa lako la `.red-bg` linapaswa kuwa na `background-color` ya nyekundu.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
