---
id: 587d7dbf367417b2b2512bbc
title: Gawanya mitindo yako katika vipande vidogo kwa kutumia partials
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>Partials</dfn> katika Sass ni mafaili tofauti yanayoshikilia vipande vya msimbo wa CSS. Hivi vinaletwa na kutumika katika mafaili mengine ya Sass. Hii ni njia nzuri ya kuunganisha msimbo unaofanana katika moduli ili kuuweka umeandaliwa.

Majina ya partials huanza na herufi ya underscore (`_`) ambayo inamwambia Sass kuwa ni kipande kidogo cha CSS na isiibadilishe kuwa faili ya CSS. Pia, mafaili ya Sass huisha na kiambatisho cha faili `.scss`. Ili kuleta msimbo ulioko katika partial katika faili nyingine ya Sass, tumia amri ya `@import`.

Kwa mfano, kama mixins zako zote zimehifadhiwa katika partial iitwayo "\_mixins.scss", na zinahitajika katika faili ya "main.scss", hivi ndivyo unavyotumia katika faili kuu:

```scss
@import 'mixins'
```

Kumbuka kuwa underscore na kiambatisho cha faili havihitajiki katika tamko la `import` - Sass inaelewa kuwa ni partial. Mara partial inapoletewa katika faili, vigezo vyote, mixins, na msimbo mwingine vinapatikana kwa matumizi.

# --instructions--

Andika tamko la `@import` kuleta partial iitwayo `_variables.scss` katika faili la main.scss.

# --hints--

Msimbo wako unapaswa kutumia amri ya `@import`, na usijumuishe underscore katika jina la faili.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
