---
id: 587d7dbd367417b2b2512bb4
title: Hifadhi data kwa kutumia vigezo vya Sass
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Kipengele kimoja cha Sass kinachotofautiana na CSS ni kwamba hutumia vigezo. Hutatangazwa na kuwekwa kuhifadhi data, kama vile JavaScript.

Katika JavaScript, vigezo hutangazwa kwa kutumia maneno muhimu ya `let` na `const`. Katika Sass, vigezo huanza na `$` ikifuatiwa na jina la kigezo.

Hapa kuna mifano michache:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

Na kutumia vigezo:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

Mfano mmoja ambapo vigezo ni muhimu ni pale vipengele vingi vinapohitaji kuwa na rangi ileile. Ikiwa rangi hiyo itabadilishwa, sehemu pekee ya kuhariri msimbo ni thamani ya kigezo.

# --instructions--

Tengeneza kigezo `$text-color` na uweke thamani `red`. Kisha badilisha thamani ya sifa ya `color` kwa `.blog-post` na `h2` kuwa kigezo `$text-color`.

# --hints--

Msimbo wako unapaswa kuwa na kigezo cha Sass kilichotangazwa kwa `$text-color` na thamani ya `red`.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

Msimbo wako unapaswa kutumia kigezo `$text-color` kubadilisha sifa ya `color` kwa vitu vya `.blog-post` na `h2`.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

Kipengele chako cha `.blog-post` kinapaswa kuwa na `color` ya rangi nyekundu.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

Vipengele vyako vya `h2` vinapaswa kuwa na `color` ya rangi nyekundu.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
