---
id: 587d7fa5367417b2b2512bbd
title: Panua seti moja ya mitindo ya CSS kwa kipengele kingine
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sass ina kipengele kinachoitwa `extend` kinachorahisisha kukopa sheria za CSS kutoka kwa kipengele kimoja na kujenga juu yake katika kingine.

Kwa mfano, kipande cha chini cha sheria za CSS kinapanga darasa la `.panel`. Kina `background-color`, `height` na `border`.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

Sasa unataka paneli nyingine iitwayo `.big-panel`. Ina vigezo vya msingi sawa na `.panel`, lakini pia inahitaji `width` na `font-size`. Inawezekana kunakili na kubandika sheria za awali za CSS kutoka `.panel`, lakini msimbo unakuwa wa kurudia-rudia unapoongeza aina zaidi za paneli. Amri ya `extend` ni njia rahisi ya kutumia tena sheria zilizotungwa kwa kipengele kimoja, kisha kuongeza zaidi kwa kingine:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel` itakuwa na vigezo sawa na `.panel` pamoja na mitindo mipya.

# --instructions--

Tengeneza darasa la `.info-important` linalopanua `.info` na pia lina seti ya `background-color` iliyowekwa kuwa magenta.

# --hints--

Darasa lako la `info-important` linapaswa kuwa na `background-color` iliyowekwa kuwa `magenta`.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

Darasa lako la `info-important` linapaswa kutumia `@extend` kurithi mtindo kutoka kwa darasa la `info`.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
