---
id: 587d7dbe367417b2b2512bb8
title: Tumia @if na @else kuongeza mantiki kwa mitindo yako
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Amri ya `@if` katika Sass ni muhimu kwa kujaribu kesi maalum - inafanya kazi kama amri ya `if` katika JavaScript.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

Na kama ilivyo katika JavaScript, amri za `@else if` na `@else` hujaribu masharti zaidi:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

Tengeneza mixin iitwayo `border-stroke` inayopokea kigezo `$val`. Mixin inapaswa kuangalia masharti yafuatayo kwa kutumia amri za `@if`, `@else if`, na `@else`:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

Kama thamani ya kigezo `$val` si `light`, `medium`, au `heavy`, basi sifa ya `border` inapaswa kuwekwa kuwa `none`.

# --hints--

Msimbo wako unapaswa kutangaza mixin iitwayo `border-stroke` ambayo ina kigezo kinachoitwa `$val`.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

Mixin yako inapaswa kuwa na amri ya `@if` kuangalia kama `$val` ni `light`, na kuweka `border` kuwa `1px solid black`.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Mixin yako inapaswa kuwa na amri ya `@else if` kuangalia kama `$val` ni `medium`, na kuweka `border` kuwa `3px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Mixin yako inapaswa kuwa na amri ya `@else if` kuangalia kama `$val` ni `heavy`, na kuweka `border` kuwa `6px solid black`.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

Mixin yako inapaswa kuwa na amri ya `@else` kuweka `border` kuwa `none`.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
