---
id: 587d7dbd367417b2b2512bb5
title: Panga ndani CSS na Sass
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass inaruhusu upangaji ndani wa sheria za CSS, ambayo ni mbinu muhimu ya kupanga karatasi ya mtindo.

Kawaida, kila kipengele kinawalengwa kwenye mstari tofauti ili kukiandaa mtindo, kama ifuatavyo:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

Kwa mradi mkubwa, faili la CSS litakuwa na mistari mingi na sheria nyingi. Hapa ndipo upangaji ndani unavyoweza kusaidia kupanga msimbo wako kwa kuweka sheria za mtindo za watoto ndani ya vipengele vya mzazi husika:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

Tumia mbinu ya upangaji ndani iliyoonyeshwa hapo juu kupanga upya sheria za CSS kwa watoto wote wa kipengele cha `.blog-post`. Kwa madhumuni ya majaribio, `h1` inapaswa kuja kabla ya kipengele cha `p`.

# --hints--

Msimbo wako unapaswa kupanga upya sheria za CSS ili `h1` na `p` zipangwe ndani ya kipengele mzazi cha `.blog-post`.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
