---
id: 6823ac607bfdbc46331b2559
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTMLはHyperText Markup Languageの略です。ウェブページの構造と内容を定義するコードです。ここがコードエディターで、HTMLを書く場所です。

エディターの1行目を見つけて、次のテキストを入力してください。

```md
Welcome to freeCodeCamp
```

入力が終わったら、「コードをチェック」ボタンをクリックして正しいか確認してください。

# --hints--

エディターに`Welcome to freeCodeCamp`というテキストがあるはずです。スペルをもう一度確認してください。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
