---
id: 6823d6244511f252c8300eed
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

ウェブページに段落を追加する必要があるときは、次のように`p`要素を使うことができます。

```html
<p>This is a paragraph element.</p>
```

`Learn the skills to become a full-stack developer`を段落要素に変えてください。

# --hints--

`p`要素には開始タグが必要です。開始タグの構文は次の通りです：`<elementName>`。

```js
assert.exists(document.querySelector("p"));
```

`p`要素には終了タグが必要です。終了タグの構文は次の通りです：`</elementName>`。

```js
assert.match(code, /<\/p\s*\>/);
```

`p`要素のテキストは`Learn the skills to become a full-stack developer`であるべきです。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

`p`要素は`h2`要素の下にあるべきです。

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
