---
id: 6823e036ea4b71553558c01b
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

プレビューで各見出しが少しずつ異なって見えることに気づきましたか。段落も同様です。

次に、以下のテキストを表示する別の`p`要素を追加してください。

```md
HTML represents the content and structure of a webpage
```

# --hints--

2つ目の`p`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

`p`要素には`HTML represents the content and structure of a webpage`のテキストが含まれているはずです。指示からテキストをコピーできます。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
