---
id: 6823d9ac8bdc3853df65a1ff
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

HTMLには6つの見出し要素があります：`h1`から`h6`までです。これらはウェブページのセクションの重要度を示すために使われ、`h1`が最も重要で`h6`が最も重要度が低いです。

`p`要素の下に、テキストが次のような`h3`見出しを追加してください：

```md
Introduction to HTML
```

# --hints--

`h3`要素は開く`<h3>`タグを持つべきです。

```js
assert.exists(document.querySelector("h3"));
```

`h3`要素は閉じる`</h3>`タグを持つべきです。

```js
assert.match(code, /<\/h3\s*\>/);
```

`h3`要素のテキストは`Introduction to HTML`であるべきです。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
