---
id: 6823e637c1c0ed56f781b4fc
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

最後に、JavaScript はウェブページをインタラクティブにします。ボタンをクリックしたり、フォームを提出したりしたときにページに何をするか指示できるようになります。

ワークショップの最後のステップとして、JavaScript を説明する `h3` と `p` の要素をもう一つずつページに追加してください。

まず、次のテキストを持つ `h3` の要素を追加します。

```md
Introduction to JavaScript
```

次に、その `h3` の要素の下に、次のテキストを持つ `p` の要素を追加してください。

```md
JavaScript adds interactivity to a webpage
```

# --hints--

ページに3つ目の `h3` の要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

`h3` の要素のテキストは `Introduction to JavaScript` である必要があります。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

ページに4つ目の `p` の要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

`p` の要素のテキストは `JavaScript adds interactivity to a webpage` である必要があります。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
