---
id: 682ba2318000b62f179bdf04
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

HTMLは要素で構成されています。最初に使うのは`h1`要素です。

```html
<h1>Welcome to freeCodeCamp</h1>
```

これは開始タグ（`<h1>`）で始まり、終了タグ（`</h1>`）で終わり、その間に表示するテキストが入ります。

`Welcome to freeCodeCamp`のテキストを、前に開始タグを、後ろに終了タグを追加して`h1`要素にしてください。

# --hints--

`h1`要素には開始`<h1>`タグが必要です。

```js
assert.exists(document.querySelector("h1"));
```

`h1`要素には終了`</h1>`タグが必要です。

```js
assert.match(code, /<\/h1\s*\>/);
```

`h1`要素は次のようになります：`<h1>Welcome to freeCodeCamp</h1>`。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
