---
id: 682cd20b883fc7b25eb539c6
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

他の2行のテキストの下に、次を追加してください：

```md
Learn the skills to become a full-stack developer
```

# --hints--

エディターに `Learn the skills to become a full-stack developer` というテキストがあることを確認してください。スペルをもう一度チェックしましょう。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--

--fcc-editable-region--
```
