---
id: 682cd206883fc7b25eb539c5
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

エディターで書いたHTMLがプレビューに表示されることに注目してください。このワークショップでは、カリキュラムの一部のウェブページ用のHTMLを書きます。

`h1`要素の下の空行に、次の内容を入力してください。

```md
Full-Stack Curriculum
```

# --hints--

エディターに`Full-Stack Curriculum`というテキストがあるはずです。スペルをもう一度確認してください。

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
