---
id: 683921c4769fd23dbadec2fe
title: ステップ 10
challengeType: 0
dashedName: step-10
---

# --description--

HTMLがウェブページの構造と内容を定義するのに対し、CSSは色、フォント、間隔、レイアウトなどのスタイルを追加するために使われます。

先ほど追加した`h3`の下に、次のテキストを含む別の段落要素を追加してください。

```md
CSS is used to style a webpage
```

# --hints--

ページに2つ目の`h3`要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

あなたの`h3`要素のテキストは`Introduction to CSS`であるべきです。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

ページに3つ目の`p`要素を追加してください。

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

あなたの`p`要素は`CSS is used to style a webpage`というテキストであるべきです。指示からテキストをコピーしてもかまいません。

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
