---
id: 587d78b2367417b2b2512b0e
title: Elemente mit push() und unshift() zu einem Array hinzufügen
challengeType: 1
forumTopicId: 301151
dashedName: add-items-to-an-array-with-push-and-unshift
---

# --description--

Die Länge eines Arrays ist, ebenso wie die Datentypen, die es enthalten kann, nicht festgelegt. Arrays können mit einer beliebigen Anzahl von Elementen definiert werden, und Elemente können im Laufe der Zeit hinzugefügt oder entfernt werden; mit anderen Worten, Arrays sind <dfn>veränderbar</dfn>. In dieser Aufgabe betrachten wir zwei Methoden, mit denen wir ein Array programmatisch verändern können: `Array.push()` und `Array.unshift()`.

Beide Methoden nehmen ein oder mehrere Elemente als Parameter und fügen diese Elemente dem Array hinzu, auf dem die Methode aufgerufen wird; die `push()`-Methode fügt Elemente ans Ende eines Arrays hinzu, und `unshift()` fügt Elemente am Anfang hinzu. Betrachten Sie Folgendes:

```js
let twentyThree = 'XXIII';
let romanNumerals = ['XXI', 'XXII'];

romanNumerals.unshift('XIX', 'XX');
```

`romanNumerals` hätte den Wert `['XIX', 'XX', 'XXI', 'XXII']`.

```js
romanNumerals.push(twentyThree);
```

`romanNumerals` hätte den Wert `['XIX', 'XX', 'XXI', 'XXII', 'XXIII']`. Beachten Sie, dass wir auch Variablen übergeben können, was uns noch mehr Flexibilität bei der dynamischen Änderung der Daten unseres Arrays ermöglicht.

# --instructions--

Wir haben eine Funktion `mixedNumbers` definiert, der wir ein Array als Argument übergeben. Ändern Sie die Funktion, indem Sie `push()` und `unshift()` verwenden, um `'I', 2, 'three'` am Anfang des Arrays und `7, 'VIII', 9` am Ende hinzuzufügen, sodass das zurückgegebene Array die Darstellungen der Zahlen 1–9 in der richtigen Reihenfolge enthält.

# --hints--

`mixedNumbers(["IV", 5, "six"])` sollte jetzt `["I", 2, "three", "IV", 5, "six", 7, "VIII", 9]` zurückgeben

```js
assert.deepEqual(mixedNumbers(['IV', 5, 'six']), [
  'I',
  2,
  'three',
  'IV',
  5,
  'six',
  7,
  'VIII',
  9
]);
```

Die `mixedNumbers`-Funktion sollte die `push()`-Methode verwenden

```js
assert(mixedNumbers.toString().match(/\.push/));
```

Die `mixedNumbers`-Funktion sollte die `unshift()`-Methode verwenden

```js
assert(mixedNumbers.toString().match(/\.unshift/));
```

# --seed--

## --seed-contents--

```js
function mixedNumbers(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(mixedNumbers(['IV', 5, 'six']));
```

# --solutions--

```js
function mixedNumbers(arr) {
  arr.push(7,'VIII',9);
  arr.unshift('I',2,'three');
  return arr;
}
```
