---
id: 587d7b7c367417b2b2512b1a
title: Zugriff auf Eigenschaftsnamen mit Klammernotation
challengeType: 1
forumTopicId: 301150
dashedName: access-property-names-with-bracket-notation
---

# --description--

In der ersten Objekt-Aufgabe haben wir die Verwendung der Klammernotation erwähnt, um auf Eigenschaftswerte zuzugreifen, indem eine Variable ausgewertet wird. Stellen Sie sich zum Beispiel vor, dass unser `foods`-Objekt in einem Programm für eine Supermarkt-Kasse verwendet wird. Wir haben eine Funktion, die `selectedFood` festlegt, und wir möchten in unserem `foods`-Objekt nach dem Vorhandensein dieses Lebensmittels suchen. Das könnte so aussehen:

```js
let selectedFood = getCurrentFood(scannedItem);
let inventory = foods[selectedFood];
```

Dieser Code wertet den in der `selectedFood`-Variablen gespeicherten Wert aus und gibt den Wert dieses Schlüssels im `foods`-Objekt zurück oder `undefined`, falls er nicht vorhanden ist. Die Klammernotation ist sehr nützlich, weil manchmal die Objekteigenschaften erst zur Laufzeit bekannt sind oder wir dynamischer auf sie zugreifen müssen.

# --instructions--

Wir haben eine Funktion `checkInventory` definiert, die ein gescanntes Element als Argument erhält. Geben Sie den aktuellen Wert des `scannedItem`-Schlüssels im `foods`-Objekt zurück. Sie können davon ausgehen, dass nur gültige Schlüssel als Argument an `checkInventory` übergeben werden.

# --hints--

`checkInventory` sollte eine Funktion sein.

```js
assert.strictEqual(typeof checkInventory, 'function');
```

Das `foods`-Objekt sollte nur die folgenden Schlüssel-Wert-Paare enthalten: `apples: 25`, `oranges: 32`, `plums: 28`, `bananas: 13`, `grapes: 35`, `strawberries: 27`.

```js
assert.deepEqual(foods, {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
});
```

`checkInventory("apples")` sollte `25` zurückgeben.

```js
assert.strictEqual(checkInventory('apples'), 25);
```

`checkInventory("bananas")` sollte `13` zurückgeben.

```js
assert.strictEqual(checkInventory('bananas'), 13);
```

`checkInventory("strawberries")` sollte `27` zurückgeben.

```js
assert.strictEqual(checkInventory('strawberries'), 27);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(checkInventory("apples"));
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

function checkInventory(scannedItem) {
  return foods[scannedItem];
}
```
