---
id: 587d7b7d367417b2b2512b1d
title: Durchlaufen der Schlüssel eines Objekts mit einer for...in-Anweisung
challengeType: 1
forumTopicId: 301162
dashedName: iterate-through-the-keys-of-an-object-with-a-for---in-statement
---

# --description--

Manchmal müssen Sie alle Schlüssel innerhalb eines Objekts durchlaufen. Sie können dafür eine <dfn>for...in</dfn>-Schleife verwenden. Die for...in-Schleife sieht folgendermaßen aus:

```javascript
const refrigerator = {
  'milk': 1,
  'eggs': 12,
};

for (const food in refrigerator) {
  console.log(food, refrigerator[food]);
}
```

Dieser Code gibt `milk 1` und `eggs 12` aus, wobei jedes Schlüssel-Wert-Paar in einer eigenen Zeile steht.

Wir haben die Variable `food` im Schleifenkopf definiert, und diese Variable wurde bei jeder Iteration auf jeden Schlüssel des Objekts gesetzt, sodass der Name jedes Lebensmittels in der Konsole ausgegeben wurde.

**HINWEIS:** Objekte behalten keine Reihenfolge der gespeicherten Schlüssel wie Arrays bei; daher ist die Position eines Schlüssels in einem Objekt oder die relative Reihenfolge, in der er erscheint, beim Referenzieren oder Zugreifen auf diesen Schlüssel unerheblich.

# --instructions--

Wir haben eine Funktion `countOnline` definiert, die ein Argument `allUsers` akzeptiert. Verwenden Sie eine <dfn>for...in</dfn>-Anweisung innerhalb dieser Funktion, um das Objekt `allUsers` zu durchlaufen und die Anzahl der Benutzer zurückzugeben, deren `online`-Eigenschaft auf `true` gesetzt ist. Ein Beispiel für ein Objekt, das an `countOnline` übergeben werden könnte, ist unten gezeigt. Jeder Benutzer hat eine `online`-Eigenschaft, die entweder auf `true` oder `false` gesetzt ist.

```js
{
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
```

# --hints--

Die Funktion `countOnline` sollte eine `for in`-Anweisung verwenden, um die Objektschlüssel des übergebenen Objekts zu durchlaufen.

```js
assert.match(
  __helpers.removeJSComments(code),
    /for\s*\(\s*(var|let|const)\s+[a-zA-Z_$]\w*\s+in\s+[a-zA-Z_$]\w*\s*\)/
);
```

Die Funktion `countOnline` sollte `1` zurückgeben, wenn das Objekt `{ Alan: { online: false }, Jeff: { online: true }, Sarah: { online: false } }` übergeben wird.

```js
const usersObj1 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj1), 1);
```

Die Funktion `countOnline` sollte `2` zurückgeben, wenn das Objekt `{ Alan: { online: true }, Jeff: { online: false }, Sarah: { online: true } }` übergeben wird.

```js
const usersObj2 = {
  Alan: {
    online: true
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: true
  }
}
assert.equal(countOnline(usersObj2), 2);
```

Die Funktion `countOnline` sollte `0` zurückgeben, wenn das Objekt `{ Alan: { online: false }, Jeff: { online: false }, Sarah: { online: false } }` übergeben wird.

```js
const usersObj3 = {
  Alan: {
    online: false
  },
  Jeff: {
    online: false
  },
  Sarah: {
    online: false
  }
}
assert.equal(countOnline(usersObj3), 0);
```

# --seed--

## --seed-contents--

```js
const users = {
  Alan: {
    online: false
  },
  Jeff: {
    online: true
  },
  Sarah: {
    online: false
  }
}

function countOnline(allUsers) {
  // Only change code below this line

  // Only change code above this line
}

console.log(countOnline(users));
```

# --solutions--

```js
function countOnline(allUsers) {
  let numOnline = 0;
  for(const user in allUsers){
    if(allUsers[user].online) {
      numOnline++;
    }
  }
  return numOnline;
}
```
