---
id: 587d7b7c367417b2b2512b1b
title: "Das Schlüsselwort delete verwenden, um Objekteigenschaften zu entfernen"
challengeType: 1
forumTopicId: 301168
dashedName: use-the-delete-keyword-to-remove-object-properties
---

# --description--

Jetzt wissen Sie, was Objekte sind und welche grundlegenden Funktionen und Vorteile sie bieten. Kurz gesagt, sie sind Schlüssel-Wert-Speicher, die eine flexible, intuitive Möglichkeit bieten, Daten zu strukturieren, ***und*** sie ermöglichen sehr schnelle Suchzeiten. Im Verlauf der weiteren Aufgaben werden wir verschiedene gängige Operationen an Objekten beschreiben, damit Sie sich daran gewöhnen, diese nützlichen Datenstrukturen in Ihren Programmen anzuwenden.

In früheren Aufgaben haben wir Schlüssel-Wert-Paare eines Objekts sowohl hinzugefügt als auch geändert. Hier sehen wir, wie wir ein Schlüssel-Wert-Paar aus einem Objekt *entfernen* können.

Schauen wir uns unser `foods`-Objektbeispiel noch einmal an. Wenn wir den `apples`-Schlüssel entfernen wollten, können wir ihn mit dem Schlüsselwort `delete` so löschen:

```js
delete foods.apples;
```

# --instructions--

Verwenden Sie das Schlüsselwort delete, um die Schlüssel `oranges`, `plums` und `strawberries` aus dem `foods`-Objekt zu entfernen.

# --hints--

Das `foods`-Objekt sollte nur noch drei Schlüssel haben: `apples`, `grapes` und `bananas`.

```js
assert(
  !foods.hasOwnProperty('oranges') &&
    !foods.hasOwnProperty('plums') &&
    !foods.hasOwnProperty('strawberries') &&
    Object.keys(foods).length === 3
);
```

Die Schlüssel `oranges`, `plums` und `strawberries` sollten mit `delete` entfernt werden.

```js
assert(
  __helpers.removeJSComments(code).search(/oranges:/) !== -1 &&
    __helpers.removeJSComments(code).search(/plums:/) !== -1 &&
    __helpers.removeJSComments(code).search(/strawberries:/) !== -1
);
```

# --seed--

## --seed-contents--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

// Only change code below this line

// Only change code above this line

console.log(foods);
```

# --solutions--

```js
let foods = {
  apples: 25,
  oranges: 32,
  plums: 28,
  bananas: 13,
  grapes: 35,
  strawberries: 27
};

delete foods.oranges;
delete foods.plums;
delete foods.strawberries;

console.log(foods);
```
