---
id: 587d7b7b367417b2b2512b13
title: Ein Array mit dem Spread-Operator kopieren
challengeType: 1
forumTopicId: 301157
dashedName: copy-an-array-with-the-spread-operator
---

# --description--

Während `slice()` uns erlaubt, selektiv auszuwählen, welche Elemente eines Arrays kopiert werden sollen, ermöglicht uns der neue <dfn>Spread-Operator</dfn> von ES6, *alle* Elemente eines Arrays einfach und in der richtigen Reihenfolge mit einer einfachen und sehr gut lesbaren Syntax zu kopieren. Die Spread-Syntax sieht einfach so aus: `...`

In der Praxis können wir den Spread-Operator verwenden, um ein Array so zu kopieren:

```js
let thisArray = [true, true, undefined, false, null];
let thatArray = [...thisArray];
```

`thatArray` ist gleich `[true, true, undefined, false, null]`. `thisArray` bleibt unverändert und `thatArray` enthält dieselben Elemente wie `thisArray`.

# --instructions--

Wir haben eine Funktion `copyMachine` definiert, die `arr` (ein Array) und `num` (eine Zahl) als Argumente entgegennimmt. Die Funktion soll ein neues Array zurückgeben, das aus `num` Kopien von `arr` besteht. Wir haben den Großteil der Arbeit bereits erledigt, aber es funktioniert noch nicht ganz richtig. Ändern Sie die Funktion mit Spread-Syntax, damit sie korrekt funktioniert (Hinweis: Eine andere Methode, die wir bereits behandelt haben, könnte hier hilfreich sein!).

# --hints--

`copyMachine([true, false, true], 2)` sollte `[[true, false, true], [true, false, true]]` zurückgeben

```js
assert.deepEqual(copyMachine([true, false, true], 2), [
  [true, false, true],
  [true, false, true]
]);
```

`copyMachine([1, 2, 3], 5)` sollte `[[1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3], [1, 2, 3]]` zurückgeben

```js
assert.deepEqual(copyMachine([1, 2, 3], 5), [
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3],
  [1, 2, 3]
]);
```

`copyMachine([true, true, null], 1)` sollte `[[true, true, null]]` zurückgeben

```js
assert.deepEqual(copyMachine([true, true, null], 1), [[true, true, null]]);
```

`copyMachine(["it works"], 3)` sollte `[["it works"], ["it works"], ["it works"]]` zurückgeben

```js
assert.deepEqual(copyMachine(['it works'], 3), [
  ['it works'],
  ['it works'],
  ['it works']
]);
```

Die `copyMachine`-Funktion sollte den `spread operator` mit dem Array `arr` verwenden

```js
assert(__helpers.removeJSComments(__helpers.removeJSComments(code)).match(/\.\.\.\s*arr/));
```

# --seed--

## --seed-contents--

```js
function copyMachine(arr, num) {
  let newArr = [];
  while (num >= 1) {
    // Only change code below this line

    // Only change code above this line
    num--;
  }
  return newArr;
}

console.log(copyMachine([true, false, true], 2));
```

# --solutions--

```js
function copyMachine(arr,num){
    let newArr=[];
    while(num >=1){
    newArr.push([...arr]);
    num--;
    }
    return newArr;
}
console.log(copyMachine([true, false, true], 2));
```
