---
id: 587d7b7b367417b2b2512b14
title: "Prüfen, ob ein Element mit indexOf() vorhanden ist"
challengeType: 1
forumTopicId: 301154
dashedName: check-for-the-presence-of-an-element-with-indexof
---

# --description--

Da Arrays jederzeit verändert oder *mutiert* werden können, gibt es keine Garantie dafür, wo sich ein bestimmtes Datum in einem Array befindet oder ob dieses Element überhaupt noch existiert. Glücklicherweise stellt JavaScript uns eine weitere eingebaute Methode, `indexOf()`, zur Verfügung, mit der wir schnell und einfach prüfen können, ob ein Element in einem Array vorhanden ist. `indexOf()` nimmt ein Element als Parameter und gibt beim Aufrufen die Position oder den Index dieses Elements zurück oder `-1`, falls das Element nicht im Array vorhanden ist.

Zum Beispiel:

```js
let fruits = ['apples', 'pears', 'oranges', 'peaches', 'pears'];

fruits.indexOf('dates');
fruits.indexOf('oranges');
fruits.indexOf('pears');
```

`indexOf('dates')` gibt `-1` zurück, `indexOf('oranges')` gibt `2` zurück und `indexOf('pears')` gibt `1` zurück (der erste Index, an dem jedes Element existiert).

# --instructions--

`indexOf()` kann sehr nützlich sein, um schnell zu prüfen, ob ein Element in einem Array vorhanden ist. Wir haben eine Funktion `quickCheck` definiert, die ein Array und ein Element als Argumente nimmt. Ändern Sie die Funktion mit `indexOf()` so, dass sie `true` zurückgibt, wenn das übergebene Element im Array vorhanden ist, und `false`, wenn nicht.

# --hints--

Die `quickCheck`-Funktion sollte einen Boolean (`true` oder `false`) zurückgeben, nicht einen String (`"true"` oder `"false"`)

```js
assert.isBoolean(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

`quickCheck(["squash", "onions", "shallots"], "mushrooms")` sollte `false` zurückgeben

```js
assert.strictEqual(
  quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'),
  false
);
```

`quickCheck(["onions", "squash", "shallots"], "onions")` sollte `true` zurückgeben

```js
assert.strictEqual(
  quickCheck(['onions', 'squash', 'shallots'], 'onions'),
  true
);
```

`quickCheck([3, 5, 9, 125, 45, 2], 125)` sollte `true` zurückgeben

```js
assert.strictEqual(quickCheck([3, 5, 9, 125, 45, 2], 125), true);
```

`quickCheck([true, false, false], undefined)` sollte `false` zurückgeben

```js
assert.strictEqual(quickCheck([true, false, false], undefined), false);
```

Die `quickCheck`-Funktion sollte die `indexOf()`-Methode verwenden

```js
assert.notStrictEqual(quickCheck.toString().search(/\.indexOf\(/), -1);
```

# --seed--

## --seed-contents--

```js
function quickCheck(arr, elem) {
  // Only change code below this line

  // Only change code above this line
}

console.log(quickCheck(['squash', 'onions', 'shallots'], 'mushrooms'));
```

# --solutions--

```js
function quickCheck(arr, elem) {
  return arr.indexOf(elem) >= 0; 
}
```
