---
id: 587d7b7b367417b2b2512b15
title: Durchlaufen aller Elemente eines Arrays mit for-Schleifen
challengeType: 1
forumTopicId: 301161
dashedName: iterate-through-all-an-arrays-items-using-for-loops
---

# --description--

Manchmal ist es beim Arbeiten mit Arrays sehr praktisch, jedes Element durchlaufen zu können, um ein oder mehrere Elemente zu finden, die wir benötigen, oder um ein Array basierend darauf zu manipulieren, welche Daten bestimmte Kriterien erfüllen. JavaScript bietet mehrere eingebaute Methoden, die Arrays jeweils auf leicht unterschiedliche Weise durchlaufen, um verschiedene Ergebnisse zu erzielen (wie `every()`, `forEach()`, `map()` usw.), doch die Technik, die am flexibelsten ist und uns die größte Kontrolle bietet, ist eine einfache `for`-Schleife.

Betrachten Sie Folgendes:

```js
function greaterThanTen(arr) {
  let newArr = [];
  for (let i = 0; i < arr.length; i++) {
    if (arr[i] > 10) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}

greaterThanTen([2, 12, 8, 14, 80, 0, 1]);
```

Mit einer `for`-Schleife durchläuft diese Funktion jedes Element des Arrays, greift darauf zu und unterzieht es einem einfachen Test, den wir erstellt haben. So haben wir leicht und programmatisch bestimmt, welche Daten größer als `10` sind, und ein neues Array, `[12, 14, 80]`, zurückgegeben, das diese Elemente enthält.

# --instructions--

Wir haben eine Funktion `filteredArray` definiert, die `arr`, ein verschachteltes Array, und `elem` als Argumente nimmt und ein neues Array zurückgibt. `elem` stellt ein Element dar, das in einem oder mehreren der Arrays, die in `arr` verschachtelt sind, vorhanden sein kann oder auch nicht. Ändern Sie die Funktion so, dass sie mit einer `for`-Schleife eine gefilterte Version des übergebenen Arrays zurückgibt, bei der alle Arrays, die `arr` enthalten und in `elem` verschachtelt sind, entfernt wurden.

# --hints--

`filteredArray([[10, 8, 3], [14, 6, 23], [3, 18, 6]], 18)` sollte `[[10, 8, 3], [14, 6, 23]]` zurückgeben

```js
assert.deepEqual(
  filteredArray(
    [
      [10, 8, 3],
      [14, 6, 23],
      [3, 18, 6]
    ],
    18
  ),
  [
    [10, 8, 3],
    [14, 6, 23]
  ]
);
```

`filteredArray([["trumpets", 2], ["flutes", 4], ["saxophones", 2]], 2)` sollte `[["flutes", 4]]` zurückgeben

```js
assert.deepEqual(
  filteredArray(
    [
      ['trumpets', 2],
      ['flutes', 4],
      ['saxophones', 2]
    ],
    2
  ),
  [['flutes', 4]]
);
```

`filteredArray([["amy", "beth", "sam"], ["dave", "sean", "peter"]], "peter")` sollte `[["amy", "beth", "sam"]]` zurückgeben

```js
assert.deepEqual(
  filteredArray(
    [
      ['amy', 'beth', 'sam'],
      ['dave', 'sean', 'peter']
    ],
    'peter'
  ),
  [['amy', 'beth', 'sam']]
);
```

`filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3)` sollte `[]` zurückgeben

```js
assert.deepEqual(
  filteredArray(
    [
      [3, 2, 3],
      [1, 6, 3],
      [3, 13, 26],
      [19, 3, 9]
    ],
    3
  ),
  []
);
```

Die Funktion `filteredArray` sollte eine `for`-Schleife verwenden

```js
assert.notStrictEqual(filteredArray.toString().search(/for/), -1);
```

# --seed--

## --seed-contents--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  // Only change code below this line

  // Only change code above this line
  return newArr;
}

console.log(filteredArray([[3, 2, 3], [1, 6, 3], [3, 13, 26], [19, 3, 9]], 3));
```

# --solutions--

```js
function filteredArray(arr, elem) {
  let newArr = [];
  for (let i = 0; i<arr.length; i++) {
    if (arr[i].indexOf(elem) < 0) {
      newArr.push(arr[i]);
    }
  }
  return newArr;
}
```
