---
id: 587d7b7b367417b2b2512b17
title: Arrays mit dem Spread-Operator kombinieren
challengeType: 1
forumTopicId: 301156
dashedName: combine-arrays-with-the-spread-operator
---

# --description--

Ein weiterer großer Vorteil des <dfn>Spread</dfn>-Operators ist die Möglichkeit, Arrays zu kombinieren oder alle Elemente eines Arrays an einer beliebigen Stelle in ein anderes einzufügen. Mit traditionelleren Syntaxen können wir Arrays zwar verketten, aber das erlaubt nur, Arrays am Ende des einen und am Anfang des anderen zusammenzufügen. Die Spread-Syntax macht die folgende Operation extrem einfach:

```js
let thisArray = ['sage', 'rosemary', 'parsley', 'thyme'];

let thatArray = ['basil', 'cilantro', ...thisArray, 'coriander'];
```

`thatArray` hätte den Wert `['basil', 'cilantro', 'sage', 'rosemary', 'parsley', 'thyme', 'coriander']`.

Mit der Spread-Syntax haben wir gerade eine Operation erreicht, die mit traditionellen Methoden komplexer und ausführlicher gewesen wäre.

# --instructions--

Wir haben eine Funktion `spreadOut` definiert, die die Variable `sentence` zurückgibt. Ändern Sie die Funktion mit dem <dfn>Spread</dfn>-Operator so, dass sie das Array `['learning', 'to', 'code', 'is', 'fun']` zurückgibt.

# --hints--

`spreadOut` sollte `["learning", "to", "code", "is", "fun"]` zurückgeben

```js
assert.deepEqual(spreadOut(), ['learning', 'to', 'code', 'is', 'fun']);
```

Die Funktion `spreadOut` sollte die Spread-Syntax verwenden

```js
assert.notStrictEqual(spreadOut.toString().search(/[...]/), -1);
```

# --seed--

## --seed-contents--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence; // Change this line
  return sentence;
}

console.log(spreadOut());
```

# --solutions--

```js
function spreadOut() {
  let fragment = ['to', 'code'];
  let sentence = ['learning', ...fragment, 'is', 'fun'];
  return sentence;
}
```
