---
id: 587d7b7a367417b2b2512b12
title: Array-Elemente mit slice() kopieren
challengeType: 1
forumTopicId: 301158
dashedName: copy-array-items-using-slice
---

# --description--

Die nächste Methode, die wir behandeln, ist `slice()`. Anstatt ein Array zu verändern, kopiert oder *extrahiert* `slice()` eine bestimmte Anzahl von Elementen in ein neues Array, wobei das Array, auf dem sie aufgerufen wird, unverändert bleibt. `slice()` nimmt nur 2 Parameter — der erste ist der Index, an dem die Extraktion beginnen soll, und der zweite ist der Index, an dem die Extraktion endet (die Extraktion erfolgt bis, aber ohne das Element an diesem Index). Betrachten Sie folgendes Beispiel:

```js
let weatherConditions = ['rain', 'snow', 'sleet', 'hail', 'clear'];

let todaysWeather = weatherConditions.slice(1, 3);
```

`todaysWeather` hätte den Wert `['snow', 'sleet']`, während `weatherConditions` weiterhin `['rain', 'snow', 'sleet', 'hail', 'clear']` hätte.

Effektiv haben wir ein neues Array erstellt, indem wir Elemente aus einem bestehenden Array extrahiert haben.

# --instructions--

Wir haben eine Funktion namens `forecast` definiert, die ein Array als Argument erhält. Ändern Sie die Funktion so, dass sie mit `slice()` Informationen aus dem Argument-Array extrahiert und ein neues Array zurückgibt, das die String-Elemente `warm` und `sunny` enthält.

# --hints--

`forecast` sollte `["warm", "sunny"]` zurückgeben

```js
assert.deepEqual(
  forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']),
  ['warm', 'sunny']
);
```

Die `forecast`-Funktion sollte die `slice()`-Methode verwenden

```js
assert(/\.slice\(/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function forecast(arr) {
  // Only change code below this line

  return arr;
}

// Only change code above this line
console.log(forecast(['cold', 'rainy', 'warm', 'sunny', 'cool', 'thunderstorms']));
```

# --solutions--

```js
function forecast(arr) {
  return arr.slice(2,4);
}
```
