---
id: 587d78b3367417b2b2512b11
title: Elemente mit splice() hinzufügen
challengeType: 1
forumTopicId: 301152
dashedName: add-items-using-splice
---

# --description--

Erinnern Sie sich, dass wir in der letzten Aufgabe erwähnt haben, dass `splice()` bis zu drei Parameter annehmen kann? Nun, Sie können den dritten Parameter, der aus einem oder mehreren Elementen besteht, verwenden, um Elemente zum Array hinzuzufügen. Das ist äußerst nützlich, um schnell ein Element oder eine Gruppe von Elementen durch andere zu ersetzen.

```js
const numbers = [10, 11, 12, 12, 15];
const startIndex = 3;
const amountToDelete = 1;

numbers.splice(startIndex, amountToDelete, 13, 14);
console.log(numbers);
```

Das zweite Vorkommen von `12` wird entfernt, und wir fügen `13` und `14` an derselben Stelle ein. Das `numbers`-Array wäre jetzt `[ 10, 11, 12, 13, 14, 15 ]`.

Hier beginnen wir mit einem Array von Zahlen. Dann übergeben wir an `splice()` Folgendes: Den Index, an dem Elemente gelöscht werden sollen (3), die Anzahl der zu löschenden Elemente (1) und die übrigen Argumente (13, 14), die an genau dieser Stelle eingefügt werden. Beachten Sie, dass beliebig viele Elemente (durch Kommas getrennt) nach `amountToDelete` folgen können, die alle eingefügt werden.

# --instructions--

Wir haben eine Funktion `htmlColorNames` definiert, die ein Array von HTML-Farben als Argument erhält. Ändern Sie die Funktion mit `splice()`, um die ersten zwei Elemente des Arrays zu löschen und `'DarkSalmon'` und `'BlanchedAlmond'` an deren Stellen hinzuzufügen.

# --hints--

`htmlColorNames` sollte `["DarkSalmon", "BlanchedAlmond", "LavenderBlush", "PaleTurquoise", "FireBrick"]` zurückgeben

```js
assert.deepEqual(
  htmlColorNames([
    'DarkGoldenRod',
    'WhiteSmoke',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]),
  [
    'DarkSalmon',
    'BlanchedAlmond',
    'LavenderBlush',
    'PaleTurquoise',
    'FireBrick'
  ]
);
```

Die `htmlColorNames`-Funktion sollte die `splice()`-Methode verwenden

```js
assert(/.splice/.test(__helpers.removeJSComments(code)));
```

Sie sollten `shift()` oder `unshift()` nicht verwenden.

```js
assert(!/shift|unshift/.test(__helpers.removeJSComments(code)));
```

Sie sollten keine Klammernotation für Arrays verwenden.

```js
assert(!/\[\d\]\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function htmlColorNames(arr) {
  // Only change code below this line

  // Only change code above this line
  return arr;
}

console.log(htmlColorNames(['DarkGoldenRod', 'WhiteSmoke', 'LavenderBlush', 'PaleTurquoise', 'FireBrick']));
```

# --solutions--

```js
function htmlColorNames(arr) {
  arr.splice(0,2,'DarkSalmon', 'BlanchedAlmond');
  return arr;
}
```
