---
id: 587d78b2367417b2b2512b0f
title: Elemente aus einem Array mit pop() und shift() entfernen
challengeType: 1
forumTopicId: 301165
dashedName: remove-items-from-an-array-with-pop-and-shift
---

# --description--

Sowohl `push()` als auch `unshift()` haben entsprechende Methoden, die nahezu funktionale Gegensätze sind: `pop()` und `shift()`. Wie Sie vielleicht schon vermutet haben, entfernt `pop()` *ein* Element vom Ende eines Arrays, während `shift()` ein Element vom Anfang entfernt. Der entscheidende Unterschied zwischen `pop()` und `shift()` und ihren Pendants `push()` und `unshift()` ist, dass keine der Methoden Parameter annimmt und jeweils nur ein Element auf einmal aus dem Array entfernt werden kann.

Schauen wir uns das an:

```js
let greetings = ['whats up?', 'hello', 'see ya!'];

greetings.pop();
```

`greetings` hätte den Wert `['whats up?', 'hello']`.

```js
greetings.shift();
```

`greetings` hätte den Wert `['hello']`.

Wir können auch den Wert des entfernten Elements mit beiden Methoden so zurückgeben:

```js
let popped = greetings.pop();
```

`greetings` hätte den Wert `[]`, und `popped` hätte den Wert `hello`.

# --instructions--

Wir haben eine Funktion `popShift` definiert, die ein Array als Argument nimmt und ein neues Array zurückgibt. Ändern Sie die Funktion so, dass sie mit `pop()` und `shift()` das erste und letzte Element des Argument-Arrays entfernt und die entfernten Elemente den entsprechenden Variablen zuweist, sodass das zurückgegebene Array deren Werte enthält.

# --hints--

`popShift(["challenge", "is", "not", "complete"])` sollte `["challenge", "complete"]` zurückgeben

```js
assert.deepEqual(popShift(['challenge', 'is', 'not', 'complete']), [
  'challenge',
  'complete'
]);
```

Die `popShift`-Funktion sollte die `pop()`-Methode verwenden

```js
assert.notStrictEqual(popShift.toString().search(/\.pop\(/), -1);
```

Die `popShift`-Funktion sollte die `shift()`-Methode verwenden

```js
assert.notStrictEqual(popShift.toString().search(/\.shift\(/), -1);
```

# --seed--

## --seed-contents--

```js
function popShift(arr) {
  let popped; // Change this line
  let shifted; // Change this line
  return [shifted, popped];
}

console.log(popShift(['challenge', 'is', 'not', 'complete']));
```

# --solutions--

```js
function popShift(arr) {
  let popped = arr.pop(); // Change this line
  let shifted = arr.shift(); // Change this line
  return [shifted, popped];
}
```
