---
id: 587d7b7d367417b2b2512b1f
title: Ein Array in einem Objekt ändern
challengeType: 1
forumTopicId: 301163
dashedName: modify-an-array-stored-in-an-object
---

# --description--

Jetzt haben Sie alle grundlegenden Operationen für JavaScript-Objekte kennengelernt. Sie können Schlüssel-Wert-Paare hinzufügen, ändern und entfernen, prüfen, ob Schlüssel existieren, und alle Schlüssel in einem Objekt durchlaufen. Während Sie JavaScript weiter lernen, werden Sie noch vielseitigere Anwendungen von Objekten sehen. Außerdem behandeln die Lektionen zu Datenstrukturen im Abschnitt Coding Interview Prep des Curriculums auch die ES6-Objekte <dfn>Map</dfn> und <dfn>Set</dfn>, die beide ähnlich wie gewöhnliche Objekte sind, aber einige zusätzliche Funktionen bieten. Jetzt, da Sie die Grundlagen von Arrays und Objekten gelernt haben, sind Sie bestens vorbereitet, um komplexere Probleme mit JavaScript anzugehen!

# --instructions--

Sehen Sie sich das Objekt an, das wir im Code-Editor bereitgestellt haben. Das `user`-Objekt enthält drei Schlüssel. Der `data`-Schlüssel enthält fünf Schlüssel, von denen einer ein Array von `friends` enthält. Daraus können Sie erkennen, wie flexibel Objekte als Datenstrukturen sind. Wir haben begonnen, eine Funktion `addFriend` zu schreiben. Beenden Sie sie so, dass sie ein `user`-Objekt entgegennimmt, den Namen des `friend`-Arguments zum Array speichert, das in `user.data.friends` gespeichert ist, und dieses Array zurückgibt.

# --hints--

Das `user`-Objekt sollte die Schlüssel `name`, `age` und `data` haben.

```js
assert('name' in user && 'age' in user && 'data' in user);
```

Die `addFriend`-Funktion sollte ein `user`-Objekt und einen `friend`-String als Argumente akzeptieren und den Freund zum Array von `friends` im `user`-Objekt hinzufügen.

```js
assert(
  (function () {
    let L1 = user.data.friends.length;
    addFriend(user, 'Sean');
    let L2 = user.data.friends.length;
    return L2 === L1 + 1;
  })()
);
```

`addFriend(user, "Pete")` sollte `["Sam", "Kira", "Tomo", "Pete"]` zurückgeben.

```js
assert.deepEqual(
  (function () {
    delete user.data.friends;
    user.data.friends = ['Sam', 'Kira', 'Tomo'];
    return addFriend(user, 'Pete');
  })(),
  ['Sam', 'Kira', 'Tomo', 'Pete']
);
```

# --seed--

## --seed-contents--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  // Only change code below this line

  // Only change code above this line
}

console.log(addFriend(user, 'Pete'));
```

# --solutions--

```js
let user = {
  name: 'Kenneth',
  age: 28,
  data: {
    username: 'kennethCodesAllDay',
    joinDate: 'March 26, 2016',
    organization: 'freeCodeCamp',
    friends: [
      'Sam',
      'Kira',
      'Tomo'
    ],
    location: {
      city: 'San Francisco',
      state: 'CA',
      country: 'USA'
    }
  }
};

function addFriend(userObj, friend) {
  userObj.data.friends.push(friend);
  return userObj.data.friends;
}
```
