---
id: 587d7b7d367417b2b2512b1e
title: Ein Array aller Objektschlüssel mit Object.keys() erzeugen
challengeType: 1
forumTopicId: 301160
dashedName: generate-an-array-of-all-object-keys-with-object-keys
---

# --description--

Wir können auch ein Array erzeugen, das alle in einem Objekt gespeicherten Schlüssel enthält, mit der `Object.keys()`-Methode. Diese Methode nimmt ein Objekt als Argument und gibt ein Array von Strings zurück, die jede Eigenschaft im Objekt darstellen. Auch hier gibt es keine bestimmte Reihenfolge der Einträge im Array.

# --instructions--

Beenden Sie die `getArrayOfUsers`-Funktion so, dass sie ein Array zurückgibt, das alle Eigenschaften des Objekts enthält, das sie als Argument erhält.

# --hints--

Das `users`-Objekt sollte nur die Schlüssel `Alan`, `Jeff`, `Sarah` und `Ryan` enthalten

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

Die `getArrayOfUsers`-Funktion sollte ein Array zurückgeben, das alle Schlüssel im `users`-Objekt enthält

```js
assert(
  (function () {
    users.Sam = {};
    users.Lewis = {};
    let R = getArrayOfUsers(users);
    return (
      R.indexOf('Alan') !== -1 &&
      R.indexOf('Jeff') !== -1 &&
      R.indexOf('Sarah') !== -1 &&
      R.indexOf('Ryan') !== -1 &&
      R.indexOf('Sam') !== -1 &&
      R.indexOf('Lewis') !== -1
    );
  })() === true
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  // Only change code below this line

  // Only change code above this line
}

console.log(getArrayOfUsers(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: false
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: false
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function getArrayOfUsers(obj) {
  return Object.keys(obj);
}

console.log(getArrayOfUsers(users));
```
