---
id: 587d7b7b367417b2b2512b16
title: Komplexe mehrdimensionale Arrays erstellen
challengeType: 1
forumTopicId: 301159
dashedName: create-complex-multi-dimensional-arrays
---

# --description--

Super! Sie haben gerade eine Menge über Arrays gelernt! Dies war ein ziemlich umfassender Überblick, und es gibt noch viel mehr über die Arbeit mit Arrays zu lernen, vieles davon werden Sie in späteren Abschnitten sehen. Bevor wir uns jedoch <dfn>Objekte</dfn> ansehen, werfen wir noch einen Blick darauf, wie Arrays etwas komplexer werden können als in den bisherigen Aufgaben.

Eine der mächtigsten Funktionen, wenn man Arrays als Datenstrukturen betrachtet, ist, dass Arrays andere Arrays enthalten können oder sogar vollständig aus anderen Arrays bestehen können. Wir haben in früheren Aufgaben Arrays gesehen, die Arrays enthalten, aber recht einfache. Arrays können jedoch eine unendliche Tiefe von Arrays enthalten, die wiederum andere Arrays enthalten können, jeweils mit beliebigen Tiefenstufen und so weiter. Auf diese Weise kann ein Array sehr schnell zu einer sehr komplexen Datenstruktur werden, die als <dfn>mehrdimensionales</dfn> oder verschachteltes Array bekannt ist. Betrachten Sie folgendes Beispiel:

```js
let nestedArray = [
  ['deep'],
  [
    ['deeper'], ['deeper'] 
  ],
  [
    [
      ['deepest'], ['deepest']
    ],
    [
      [
        ['deepest-est?']
      ]
    ]
  ]
];
```

Das `deep`-Array ist 2 Ebenen tief verschachtelt. Die `deeper`-Arrays sind 3 Ebenen tief. Die `deepest`-Arrays sind 4 Ebenen tief, und das `deepest-est?` ist 5 Ebenen tief.

Obwohl dieses Beispiel verworren erscheinen mag, ist dieses Maß an Komplexität bei der Arbeit mit großen Datenmengen nicht ungewöhnlich oder selten. Wir können jedoch immer noch sehr einfach auf die tiefsten Ebenen eines so komplexen Arrays mit der Klammernotation zugreifen:

```js
console.log(nestedArray[2][1][0][0][0]);
```

Dies gibt den String `deepest-est?` aus. Und jetzt, da wir wissen, wo sich dieses Datenstück befindet, können wir es bei Bedarf zurücksetzen:

```js
nestedArray[2][1][0][0][0] = 'deeper still';

console.log(nestedArray[2][1][0][0][0]);
```

Jetzt gibt es `deeper still` aus.

# --instructions--

Wir haben eine Variable `myNestedArray` definiert, die einem Array zugewiesen ist. Ändern Sie `myNestedArray` so, dass es genau fünf Ebenen tief ist (denken Sie daran, das äußerste Array ist Ebene 1), und verwenden Sie dabei eine beliebige Kombination aus <dfn>Strings</dfn>, <dfn>Zahlen</dfn> und <dfn>Booleans</dfn> als Datenelemente. Fügen Sie irgendwo auf der dritten Ebene den String `deep` ein, auf der vierten Ebene den String `deeper` und auf der fünften Ebene den String `deepest`.

# --hints--

`myNestedArray` sollte nur Zahlen, Booleans und Strings als Datenelemente enthalten

```js
assert.strictEqual(
  (function (arr) {
    let flattened = (function flatten(arr) {
      const flat = [].concat(...arr);
      return flat.some(Array.isArray) ? flatten(flat) : flat;
    })(arr);
    for (let i = 0; i < flattened.length; i++) {
      if (
        typeof flattened[i] !== 'number' &&
        typeof flattened[i] !== 'string' &&
        typeof flattened[i] !== 'boolean'
      ) {
        return false;
      }
    }
    return true;
  })(myNestedArray),
  true
);
```

`myNestedArray` sollte genau 5 Ebenen tief sein

```js
assert.strictEqual(
  (function (arr) {
    let depth = 0;
    function arrayDepth(array, i, d) {
      if (Array.isArray(array[i])) {
        arrayDepth(array[i], 0, d + 1);
      } else {
        depth = d > depth ? d : depth;
      }
      if (i < array.length) {
        arrayDepth(array, i + 1, d);
      }
    }
    arrayDepth(arr, 0, 0);
    return depth;
  })(myNestedArray),
  4
);
```

`myNestedArray` sollte genau eine Vorkommnis des Strings `deep` in einem Array enthalten, das 3 Ebenen tief verschachtelt ist

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deep').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deep')[0] === 2
);
```

`myNestedArray` sollte genau eine Vorkommnis des Strings `deeper` in einem Array enthalten, das 4 Ebenen tief verschachtelt ist

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deeper').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deeper')[0] === 3
);
```

`myNestedArray` sollte genau eine Vorkommnis des Strings `deepest` in einem Array enthalten, das 5 Ebenen tief verschachtelt ist

```js
assert(
  (function howDeep(array, target, depth = 0) {
    return array.reduce((combined, current) => {
      if (Array.isArray(current)) {
        return combined.concat(howDeep(current, target, depth + 1));
      } else if (current === target) {
        return combined.concat(depth);
      } else {
        return combined;
      }
    }, []);
  })(myNestedArray, 'deepest').length === 1 &&
    (function howDeep(array, target, depth = 0) {
      return array.reduce((combined, current) => {
        if (Array.isArray(current)) {
          return combined.concat(howDeep(current, target, depth + 1));
        } else if (current === target) {
          return combined.concat(depth);
        } else {
          return combined;
        }
      }, []);
    })(myNestedArray, 'deepest')[0] === 4
);
```

# --seed--

## --seed-contents--

```js
let myNestedArray = [
  // Only change code below this line
  ['unshift', false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
  // Only change code above this line
];
```

# --solutions--

```js
let myNestedArray = [
  ['unshift', ['deep', ['deeper', ['deepest']]],false, 1, 2, 3, 'complex', 'nested'],
  ['loop', 'shift', 6, 7, 1000, 'method'],
  ['concat', false, true, 'spread', 'array'],
  ['mutate', 1327.98, 'splice', 'slice', 'push'],
  ['iterate', 1.3849, 7, '8.4876', 'arbitrary', 'depth']
];
```
