---
id: 587d7b7d367417b2b2512b1c
title: "Prüfen, ob ein Objekt eine Eigenschaft hat"
challengeType: 1
forumTopicId: 301155
dashedName: check-if-an-object-has-a-property
---

# --description--

Jetzt können wir Schlüssel zu Objekten hinzufügen, ändern und entfernen. Aber was, wenn wir nur wissen wollen, ob ein Objekt eine bestimmte Eigenschaft hat? JavaScript bietet uns zwei verschiedene Möglichkeiten dafür. Eine verwendet die `hasOwnProperty()`-Methode und die andere das Schlüsselwort `in`. Wenn wir ein Objekt `users` mit einer Eigenschaft `Alan` haben, könnten wir auf eine der folgenden Arten prüfen, ob diese vorhanden ist:

```js
users.hasOwnProperty('Alan');
'Alan' in users;
```

Beide würden `true` zurückgeben.

# --instructions--

Beenden Sie die Funktion so, dass sie `true` zurückgibt, wenn das übergebene Objekt alle vier Namen `Alan`, `Jeff`, `Sarah` und `Ryan` enthält, und andernfalls `false`.

# --hints--

Das `users`-Objekt sollte nicht direkt aufgerufen werden

```js 

assert(__helpers.removeJSComments(code).match(/users/gm).length <= 2)

```

Das `users`-Objekt sollte nur die Schlüssel `Alan`, `Jeff`, `Sarah` und `Ryan` enthalten

```js
assert(
  'Alan' in users &&
    'Jeff' in users &&
    'Sarah' in users &&
    'Ryan' in users &&
    Object.keys(users).length === 4
);
```

Die Funktion `isEveryoneHere` sollte `true` zurückgeben, wenn `Alan`, `Jeff`, `Sarah` und `Ryan` Eigenschaften des übergebenen Objekts sind.

```js
assert(isEveryoneHere(users) === true);
```

Die Funktion `isEveryoneHere` sollte `false` zurückgeben, wenn `Alan` keine Eigenschaft des übergebenen Objekts ist.

```js
assert(
  (function () {
    delete users.Alan;
    return isEveryoneHere(users);
  })() === false
);
```

Die Funktion `isEveryoneHere` sollte `false` zurückgeben, wenn `Jeff` keine Eigenschaft des übergebenen Objekts ist.

```js
assert(
  (function () {
    delete users.Jeff;
    return isEveryoneHere(users);
  })() === false
);
```

Die Funktion `isEveryoneHere` sollte `false` zurückgeben, wenn `Sarah` keine Eigenschaft des übergebenen Objekts ist.

```js
assert(
  (function () {
    delete users.Sarah;
    return isEveryoneHere(users);
  })() === false
);
```

Die Funktion `isEveryoneHere` sollte `false` zurückgeben, wenn `Ryan` keine Eigenschaft des übergebenen Objekts ist.

```js
assert(
  (function () {
    delete users.Ryan;
    return isEveryoneHere(users);
  })() === false
);
```

# --seed--

## --seed-contents--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  // Only change code below this line
  
  // Only change code above this line
}

console.log(isEveryoneHere(users));
```

# --solutions--

```js
let users = {
  Alan: {
    age: 27,
    online: true
  },
  Jeff: {
    age: 32,
    online: true
  },
  Sarah: {
    age: 48,
    online: true
  },
  Ryan: {
    age: 19,
    online: true
  }
};

function isEveryoneHere(userObj) {
  return [
    'Alan',
    'Jeff',
    'Sarah',
    'Ryan'
  ].every(user => userObj.hasOwnProperty(user));
}

console.log(isEveryoneHere(users));
```
