---
id: 5a661e0f1068aca922b3ef17
title: Zugriff auf den Inhalt eines Arrays mit Klammernotation
challengeType: 1
forumTopicId: 301149
dashedName: access-an-arrays-contents-using-bracket-notation
---

# --description--

Die grundlegende Funktion jeder Datenstruktur ist natürlich die Fähigkeit, nicht nur Daten zu speichern, sondern diese Daten auch auf Abruf wieder abzurufen. Da wir nun gelernt haben, wie man ein Array erstellt, wollen wir anfangen zu überlegen, wie wir auf die Informationen dieses Arrays zugreifen können.

Wenn wir ein einfaches Array wie unten definiert haben, enthält es 3 Elemente:

```js
let ourArray = ["a", "b", "c"];
```

In einem Array hat jedes Element einen <dfn>Index</dfn>. Dieser Index entspricht gleichzeitig der Position des Elements im Array und wie Sie darauf zugreifen. Es ist jedoch wichtig zu beachten, dass JavaScript-Arrays <dfn>nullbasiert</dfn> sind, was bedeutet, dass das erste Element eines Arrays tatsächlich an der ***nullten*** Position steht, nicht an der ersten. Um ein Element aus einem Array abzurufen, können wir einen Index in Klammern setzen und ihn an das Ende eines Arrays oder, häufiger, an eine Variable anhängen, die auf ein Array-Objekt verweist. Dies wird als <dfn>Klammernotation</dfn> bezeichnet. Wenn wir zum Beispiel das `a` aus `ourArray` abrufen und einer Variablen zuweisen wollen, können wir das mit folgendem Code tun:

```js
let ourVariable = ourArray[0];
```

Jetzt hat `ourVariable` den Wert von `a`.

Neben dem Zugriff auf den Wert, der mit einem Index verknüpft ist, können Sie mit derselben Notation auch einen Index *setzen*:

```js
ourArray[1] = "not b anymore";
```

Mit der Klammernotation haben wir nun das Element an Index 1 vom String `b` auf `not b anymore` zurückgesetzt. Jetzt ist `ourArray` `["a", "not b anymore", "c"]`.

# --instructions--

Um diese Aufgabe zu lösen, setzen Sie die 2. Position (Index `1`) von `myArray` auf einen beliebigen Wert, außer dem Buchstaben `b`.

# --hints--

`myArray[0]` sollte gleich dem Buchstaben `a` sein

```js
assert.strictEqual(myArray[0], 'a');
```

`myArray[1]` sollte nicht gleich dem Buchstaben `b` sein

```js
assert.notStrictEqual(myArray[1], 'b');
```

`myArray[2]` sollte gleich dem Buchstaben `c` sein

```js
assert.strictEqual(myArray[2], 'c');
```

`myArray[3]` sollte gleich dem Buchstaben `d` sein

```js
assert.strictEqual(myArray[3], 'd');
```

# --seed--

## --seed-contents--

```js
let myArray = ["a", "b", "c", "d"];
// Only change code below this line

// Only change code above this line
console.log(myArray);
```

# --solutions--

```js
let myArray = ["a", "b", "c", "d"];
myArray[1] = "e";
```
