---
id: 69461e4f98594012939bc1c6
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utafanya mazoezi ya aina za data za Python kwa kutengeneza chapisho rahisi la karatasi ya matokeo.

Kama ulivyojifunza katika mafundisho yaliyopita, vigezo hupewa thamani kwa njia hii:

```py
greeting = 'Hello World'
```

Katika mfano hapo juu, kigezo `greeting` kimepewa thamani ya mfuatano wa herufi, ambao ni mfululizo wa herufi uliozungukwa na nukuu.

Tengeneza kigezo kinachoitwa `name` na uweke thamani ya mfuatano wa herufi `Alice`. Kumbuka kuzunguka thamani hiyo kwa nukuu moja au alama za nukuu maradufu kama ilivyoonyeshwa katika mfano.

**Muhimu:** Python hutumia nafasi za kuingiza (spaces mwanzoni mwa mstari) kupanga msimbo. Kwa warsha hii, hakikisha hakuna nafasi za ziada mwanzoni mwa kila mstari wa msimbo. Kuongeza nafasi za ziada kutasababisha `IndentationError` na kuzuia msimbo wako kuendeshwa.

# --hints--

Unapaswa kuwa na kigezo `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

Kigezo `name` kinapaswa kuhifadhi thamani `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
