---
id: 694648acde178bb8202d9514
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Sasa unahitaji kuongeza umri wa mwanafunzi kwenye karatasi ya ripoti. Kwa hiyo utatumia nambari kamili, mojawapo ya aina za data za nambari katika Python.

Tangaza kigezo kinachoitwa `age` na ukaeze thamani ya nambari kamili `20`.

Kisha, chapisha thamani na aina ya data ya `age` pamoja, zikitenganishwa na koma. Angalia matokeo kwenye terminali yanayoonyesha thamani ya `age`, na aina yake kama `int` (nambari kamili).

# --hints--

Unapaswa kuwa na kigezo `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

Kigezo `age` kinapaswa kuhifadhi thamani `20`. Usizungushe thamani hiyo na nukuu.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Unapaswa kuchapisha `age` na `type(age)` ukitumia koma kutenganisha.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
