---
id: 694648abde178bb8202d9511
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa unapaswa kuona jina la mwanafunzi likichapishwa kwenye terminali.

Python hutoa kitendakazi kinachoitwa `type()` ambacho unaweza kutumia kukagua aina ya thamani.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

Katika mfano ulio juu, matokeo `<class 'str'>` yanamaanisha kuwa kigezo kilichopitishwa kwa kitendakazi cha `type()` ni mfuatano wa herufi.

Tumia kitendakazi cha `type()` kwa `name` kama hoja yake na chapisha matokeo kama katika mfano. Kagua matokeo kwenye terminali yanayoonyesha `name` ni ya aina ya `str` (mfuatano wa herufi).

# --hints--

Unapaswa kuchapisha matokeo ya wito wa `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
