---
id: 694648acde178bb8202d9516
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Matokeo ni `False`, ambayo yanaonyesha kuwa `score` si `int`.

Aina nyingine ya kawaida ya nambari katika Python ni `float`, ambayo inawakilisha nambari yenye desimali. Badilisha `int` na `float` katika wito wa `isinstance()` uliopo kuthibitisha hili.

# --hints--

Unapaswa kubadilisha `int` na `float` katika wito wa `isinstance(score, int)` uliopo.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
