---
id: 694648acde178bb8202d9515
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Sasa, ongeza alama ya mwanafunzi.

Tangaza kigezo kinachoitwa `score` na weka thamani `80.5`.

Ingawa `age` na `score` zote ni nambari, huenda si aina ile ile. Python hutoa kitendakazi kinachoitwa `isinstance()` ili kuangalia hili.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Tumia `isinstance()` kuangalia kama `score` ni `int`, na chapisha matokeo kwenye terminal kama ilivyoonyeshwa kwenye mfano hapo juu.

# --hints--

Unapaswa kuwa na kigezo `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

Kigezo `score` kinapaswa kuhifadhi thamani `80.5`. Usizungushe thamani hiyo kwa nukuu.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Unapaswa kuchapisha matokeo ya `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
