---
id: 694648acde178bb8202d9513
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Jina la mwanafunzi linapaswa kufuata muundo uleule kama maelezo mengine.

Ondoa matokeo ya awali ya variable ya `name`. Kisha, chapisha `name` na `type(name)` pamoja kwenye mstari mmoja zikitenganishwa na koma kama katika hatua iliyopita.

# --hints--

Haupaswi kuwa na `print(name)` katika msimbo wako.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Haupaswi kuwa na `print(type(name))` katika msimbo wako.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Unapaswa kuchapisha `name` na `type(name)` katika wito mmoja wa `print()` ukitumia separator ya koma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
