---
id: 694648acde178bb8202d9517
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Matokeo ni `True`, yakithibitisha kuwa `score` ni `float`.

Kamilisha karatasi ya ripoti kwa kuchapisha thamani ya `score` pamoja na aina yake ya data kwa kutumia tamko moja la `print()`.

# --hints--

Unapaswa kuchapisha `score` na `type(score)` kwa kutumia alama ya koma kama separator.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
