---
id: 694648acde178bb8202d9512
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Kadi ya ripoti inapaswa pia kuonyesha kama mwanafunzi yuko katika usajili kwa sasa. Hii inaweza kuwakilishwa kwa kutumia thamani ya kweli au siyo kweli.

Thamani za kweli au siyo kweli zinawakilisha hali ya ndiyo-au-la, na mara nyingi hutumika kufanya maamuzi katika msimbo. Kuna thamani mbili tu za kweli au siyo kweli: `True` na `False`.

Tangaza kigezo kinachoitwa `is_student` na weka thamani `True`.

# --hints--

Unapaswa kuwa na kigezo cha `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

Kigezo `is_student` kinapaswa kuhifadhi thamani `True`. Usizungushe thamani hiyo kwa nukuu.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
