---
id: 6947a3e516bb8a4a681274a7
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Kitendakazi cha `print()` kinaweza kuonyesha thamani zaidi ya moja kwa wakati mmoja. Tenganisha thamani kwa koma (`,`) ili kuzichapisha kwenye mstari mmoja.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Chapisha `is_student` na `type(is_student)` zote kwenye mstari mmoja ukitumia koma `,` kama ilivyoonyeshwa kwenye mfano. Kisha, angalia matokeo kwenye terminal yanayoonyesha thamani ya `is_student`, na aina yake kama `bool` (thamani ya kweli au siyo kweli).

# --hints--

Unapaswa kuchapisha `is_student` na `type(is_student)` ukitumia separator ya koma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
