---
id: 6947a0c2b6de035dd8d3eea1
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Unaweza kuchapisha thamani ya kigezo kwa kutumia kitendakazi cha `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Utajifunza zaidi kuhusu vitendakazi katika mafundisho yajayo. Kwa sasa, fahamu kwamba kitendakazi ni kipande cha msimbo kinachoweza kutumika tena ambacho kinaweza kuitwa, au kuamshwa, ili kuendesha msimbo wake, na hoja zinaweza kupitishwa kwake.

Katika mfano uliotolewa hapo juu, `print(greeting)` ni wito wa kitendakazi, na `greeting` ni hoja ya kitendakazi hicho.

Rejelea mfano na chapisha kigezo cha `name`. Angalia matokeo kwenye terminali.

# --hints--

Unapaswa kuita `print()` ukiwa na `name` kama hoja yake.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
