---
id: 6894c1b715b8d18a18ef3203
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستحصل على بعض التدريب على أصناف وكائنات Python من خلال بناء جرد للآلات الموسيقية.

لتعريف صنف في Python، استخدم الكلمة المفتاحية `class` متبوعة باسم الصنف ونقطتين.

```python
class ClassName:
    # class contents go here
```

تذكر أن أسماء الأصناف في Python عادةً ما تستخدم PascalCase، مما يعني أن كل كلمة تبدأ بحرف كبير بدون فراغات بين الكلمات.

ابدأ بإنشاء صنف يسمى `MusicalInstrument`. أضف `pass` كجسم الصنف في الوقت الحالي.

# --hints--

يجب أن تنشئ صنفًا يسمى `MusicalInstrument`.

```js
({
 test: () => assert(runPython(`_Node(_code).has_class("MusicalInstrument")`))
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
