---
id: 6895e197180c9af6988e1e19
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

أنشئ نسخة ثانية باسم `instrument_2` لآلة `Trumpet` التي هي آلة من نوع `brass`.

لاحظ كيف يمكنك إنشاء نسخ متعددة من نفس قالب الصنف. كل نسخة تخزن قيم خاصياتها بشكل مستقل عن النسخ الأخرى.

# --hints--

يجب أن تنشئ نسخة باسم `instrument_2` من صنف `MusicalInstrument`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_2, MusicalInstrument)
    `)
})
```

يجب أن تمرر `Trumpet` و`brass` كمعلمات عند إنشاء `instrument_2`.

```js
({ 
    test: () => runPython(`
        assert instrument_2.name == "Trumpet"
        assert instrument_2.instrument_type == "brass" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
--fcc-editable-region--

--fcc-editable-region--
```
