---
id: 6895e3f91c069e0068a42582
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

داخل `play` الطريقة، استخدم `print()` لعرض رسالة عن عزف الآلة الموسيقية.

استخدم سلسلة نصية بصيغة f لتضمين اسم الآلة الموسيقية في الرسالة: `The [instrument name] is fun to play!`. يجب أن يُستبدَل `[instrument name]` بالاسم الفعلي للآلة.

# --hints--

يجب أن تستخدم `play` الطريقة سلسلة نصية بصيغة f لعرض الرسالة `The [instrument name] is fun to play!`. استخدم `self.name` للوصول إلى اسم الآلة.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_call('print(f"The {self.name} is fun to play!")')`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
--fcc-editable-region--
        pass
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
