---
id: 6895e0bdb0348bf46952ab67
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

حان الوقت الآن لإنشاء نسخ من `MusicalInstrument` الخاص بك.

لإنشاء نسخة من صنف، استدعِ اسم الصنف كدالة، ومرر المعلمات المطلوبة لطريقة `__init__` (باستثناء `self` التي يتولى Python التعامل معها تلقائيًا).

```python
my_instance = ClassName(arg1, arg2)
```

عند إنشاء هذه النسخة، يستدعي Python طريقة `__init__` الخاصة بك ويخزن القيم كخصائص للنسخة تنتمي فقط لهذه النسخة المحددة.

خارج تعريف الصنف، أنشئ نسخة باسم `instrument_1` لـ `Oboe` وهو آلة من آلات النفخ الخشبية `woodwind`.

# --hints--

يجب أن تنشئ نسخة من صنف `MusicalInstrument` باسم `instrument_1`.

```js
({
    test: () => runPython(`
    assert isinstance(instrument_1, MusicalInstrument)
    `)
})
```

يجب أن تمرر `Oboe` و `woodwind` كمعلمات عند إنشاء `instrument_1`.

```js
({ 
    test: () => runPython(`
        assert instrument_1.name == "Oboe"
        assert instrument_1.instrument_type == "woodwind" 
    `) 
})
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--

--fcc-editable-region--
```
