---
id: 6895e4e023e07603bcef46f9
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

استخدم الآن الطرق التي أنشأتها.

نفذ استدعاء الطريقة `play()` على الكائن `instrument_1`. سيُطبع بذلك رسالة عن تشغيل آلة الأوبوا.

لاستدعاء طريقة على كائن، استخدم ترميز النقطة:

```python
instance_name.method_name()
```

# --hints--

يجب أن تنفذ استدعاء الطريقة `play()` على الكائن `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_stmt("instrument_1.play()")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--

--fcc-editable-region--
```
