---
id: 6895e4940e39950288e406c2
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

بعد ذلك، أضف طريقة أخرى إلى الصنف الخاص بك باسم `get_fact` تُرجع حقيقة عن الآلة الموسيقية.

يجب أن تحتوي هذه الطريقة على المعلمة `self` فقط وتُرجع سلسلة نصية بصيغة f-string تقول `The [instrument name] is part of the [instrument type] family of instruments.`

على عكس طريقة `play` التي تطبع رسالة مباشرة، يجب أن تُرجع هذه الطريقة سلسلة نصية يمكن استخدامها في مكان آخر في الشفرة الخاصة بك.

# --hints--

يجب أن تعرف طريقة باسم `get_fact` تحتوي على المعلمة `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").has_args("self")`))
});
```

يجب أن تُرجع طريقة `get_fact` سلسلة نصية بصيغة f-string تحتوي على اسم الآلة الموسيقية ونوعها.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("get_fact").find_return().is_equivalent("return f'The {self.name} is part of the {self.instrument_type} family of instruments.'")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

--fcc-editable-region--

--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
