---
id: 6895e0452b3370f26f365a46
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في الطريقة `__init__`، تحتاج إلى تعيين المعلمات إلى خصائص النسخة. خصائص النسخة هي متغيرات تنتمي إلى نسخة محددة من الصنف.

داخل طريقتك `__init__`، عيّن المعلمة `name` إلى `self.name` والمعلمة `instrument_type` إلى `self.instrument_type`.

# --hints--

يجب تعيين المعلمة `name` إلى `self.name` داخل الطريقة `__init__`.

```js
 ({
    test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[0].is_equivalent("self.name = name")`))
});
```

يجب تعيين المعلمة `instrument_type` إلى `self.instrument_type` داخل الطريقة `__init__`.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").find_body()[1].is_equivalent("self.instrument_type = instrument_type")`))
})

```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
