---
id: 6895e235935106f88281b648
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

يمكنك الوصول إلى خصائص مثيلاتك باستخدام ترميز النقطة:

```py
instance_name.attribute_name
```

هذا يسمح لك باسترجاع القيم التي تم تخزينها أثناء التهيئة.

جرّب طباعة خاصيتي `name` و `instrument_type` لكلتا المثيلتين. سترى القيم التي عيّنتها عند إنشائهما.

# --hints--

يجب طباعة خاصية `name` من `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.name)")`))
})
```

يجب طباعة خاصية `instrument_type` من `instrument_1`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.instrument_type)")`))
});
```

يجب طباعة خاصية `name` من `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.name)")`))
});
```

يجب طباعة خاصية `instrument_type` من `instrument_2`.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_2.instrument_type)")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

--fcc-editable-region--

--fcc-editable-region--
```
