---
id: 6895dffc64e97df0e62e9ef9
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

الطرق هي دوال تُعرّف داخل صنف ويمكنها تنفيذ إجراءات باستخدام بيانات الصنف. إليك كيفية إنشاء الطرق:

```python
def method_name(parameter1, parameter2,...):
    # Method contents will go here
```

في Python، تُعتبر الطريقة `__init__` طريقة خاصة تُعرف بطريقة المُهيئ، وتُستدعى عند إنشاء نسخة جديدة من الصنف.

داخل صنف `MusicalInstrument` الخاص بك، أضف طريقة `__init__` بثلاث معلمات: `self` و`name` و`instrument_type`.

المعلمة `self` هي طريقة قياسية للإشارة إلى نسخة الصنف، وتُطلب كأول معلمة لكل الطرق الخاصة بالنسخ في الصنف.

# --hints--

يجب أن تُعرّف طريقة `__init__` داخل صنف `MusicalInstrument` الخاص بك.

```js
({
 test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").has_function("__init__")`))
});
```

يجب أن تحتوي طريقة `__init__` الخاصة بك على ثلاث معلمات: `self` و`name` و`instrument_type`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("__init__").has_args("self, name, instrument_type")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
--fcc-editable-region--
    
    pass
--fcc-editable-region--
```
