---
id: 6895e5337b1914050b0af7bd
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

الآن، استدعِ الطريقة `get_fact()` على الكائن `instrument_1` واطبع النتيجة.

بما أن `get_fact()` تُرجع سلسلة نصية بدلاً من طباعتها مباشرة، يجب عليك استخدام `print()` لعرض السلسلة النصية المُرجعة.

# --hints--

يجب استدعاء `get_fact()` على `instrument_1` وطباعة النتيجة.

```js
({
  test: () => assert(runPython(`_Node(_code).has_call("print(instrument_1.get_fact())")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

    def play(self):
        print(f'The {self.name} is fun to play!')

    def get_fact(self):
        return f'The {self.name} is part of the {self.instrument_type} family of instruments.'

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')

instrument_1.play()
--fcc-editable-region--

--fcc-editable-region--
```
