---
id: 6895e356fd877efd6e3657ee
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--


داخل صنف `MusicalInstrument`، عرّف طريقة أخرى باسم `play` تحتوي فقط على معلمة `self`. تتيح معلمة `self` للطرق الوصول إلى خاصيات النسخة المحددة.

# --hints--

يجب أن تعرّف طريقة باسم `play` تحتوي على معلمة `self`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_class("MusicalInstrument").find_function("play").has_args("self")`))
});
```

# --seed--

## --seed-contents--

```py
class MusicalInstrument:
    def __init__(self, name, instrument_type):
        self.name = name
        self.instrument_type = instrument_type

--fcc-editable-region--
    
--fcc-editable-region--

instrument_1 = MusicalInstrument('Oboe', 'woodwind')
instrument_2 = MusicalInstrument('Trumpet', 'brass')
```
