---
id: 69a92f0b9cc04bb0d5327bb5
title: 实现一个栈
challengeType: 26
dashedName: implement-a-stack
---

# --description--

在本实验中，你将使用函数实现一个栈数据结构。栈是一种后进先出（LIFO）的数据结构，元素从顶部添加和移除。

**目标：** 完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `initStack` 函数，返回一个对象，该对象的 `collection` 属性是一个空数组。
1. 你应该有一个 `push` 函数，用于将元素添加到栈顶。
1. 你应该有一个 `pop` 函数，用于移除并返回栈顶元素，如果没有元素则返回 `undefined`。
1. 你应该有一个 `peek` 函数，用于返回栈顶元素但不移除它，如果没有元素则返回 `undefined`。
1. 你应该有一个 `isEmpty` 函数，如果栈中没有元素则返回 `true`，否则返回 `false`。
1. 你应该有一个 `clear` 函数，用于移除栈中所有元素。

**注意：** 大多数测试依赖于 `initStack` 和 `push` 函数。请先实现它们，因为 `pop`、`peek`、`isEmpty` 和 `clear` 的测试需要先向栈中添加元素。

# --hints--

你应该有一个 `push` 函数。

```js
assert.isFunction(push);
```

`push` 函数应该将元素添加到栈顶。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

你应该有一个 `pop` 函数。

```js
assert.isFunction(pop);
```

`pop` 函数应该移除并返回栈顶元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

`pop` 函数应该在栈为空时返回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` 函数应该正确返回假值。

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

你应该有一个 `peek` 函数。

```js
assert.isFunction(peek);
```

`peek` 函数应该返回栈顶元素但不移除它。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

`peek` 函数应该在栈为空时返回 `undefined`。

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` 函数应该正确返回假值。

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

你应该有一个 `isEmpty` 函数。

```js
assert.isFunction(isEmpty);
```

`isEmpty` 函数应该在栈为空时返回 `true`。

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` 函数应该在栈非空时返回 `false`。

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

当栈顶元素为假值时，`isEmpty` 函数应该返回 `false`。

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

你应该有一个 `clear` 函数。

```js
assert.isFunction(clear);
```

`clear` 函数应该移除栈中所有元素。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
