---
id: 698b42492a7f5f737895892e
title: Construir um Conversor de Lista de Adjacência para Matriz
challengeType: 26
dashedName: build-an-adjacency-list-to-matrix-converter
---

# --description--

Neste laboratório, você construirá uma função que converte uma representação de lista de adjacência de um grafo em uma matriz de adjacência.  An adjacency list is an object where each key represents a node, and the corresponding value is an array of nodes that the key node is connected to. Uma matriz de adjacência é um array 2D onde a entrada na posição `[i][j]` é `1` se houver uma aresta do nó `i` para o nó `j`, e `0` caso contrário.

Por exemplo, dada a lista de adjacência:

```js
{
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
}
```

A matriz de adjacência correspondente seria:

```js
[
  [0, 1, 1, 0],
  [0, 0, 1, 0],
  [1, 0, 0, 1],
  [0, 0, 1, 0]
]
```

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve definir uma `function` chamada `adjacencyListToMatrix` para converter uma lista de adjacência em uma matriz de adjacência.
2. A função deve receber um objeto que representa a lista de adjacência de um grafo não ponderado (não direcionado ou direcionado) como seu argumento.
3. A função deve:
   - Converta a lista de adjacência para uma matriz de adjacência.
   - Imprima cada linha na matriz de adjacência.
   - Retorne a matriz de adjacência.

Por exemplo, `adjacencyListToMatrix({0: [2], 1: [2, 3], 2: [0, 1, 3], 3: [1, 2]})` deve imprimir:

```md
[0, 0, 1, 0]
[0, 0, 1, 1]
[1, 1, 0, 1]
[0, 1, 1, 0]
```

e retorne `[[0, 0, 1, 0], [0, 0, 1, 1], [1, 1, 0, 1], [0, 1, 1, 0]]`.

# --hints--

Você deve definir uma função chamada `adjacencyListToMatrix`.

```js
assert.isFunction(adjacencyListToMatrix);
```

A `adjacencyListToMatrix` função deve ter um parâmetro.

```js
assert.lengthOf(adjacencyListToMatrix, 1);
```

A função deve determinar corretamente o número de nós a partir da lista de adjacência.

```js
const adjList = {0: [1], 1: [0], 2: []};
const result = adjacencyListToMatrix(adjList);
assert.lengthOf(result, 3);
result.forEach(arr => assert.lengthOf(arr, 3));
```

A função deve definir corretamente os valores da matriz para `1` para arestas existentes.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

A função deve imprimir cada linha da matriz.

```js
const spy = __helpers.spyOn(console, "log");
try {
  const adjList = {0: [1], 1: []}
  adjacencyListToMatrix(adjList);
  const args = spy.calls;
  assert.sameDeepOrderedMembers(args[0], [ [0, 1] ]);
  assert.sameDeepOrderedMembers(args[1], [ [0, 0] ]);
} catch (e) {
  assert.fail(e);
} finally {
  spy.restore();
}
```

A função deve retornar a matriz de adjacência.

```js
const adjList = {0: [1], 1: []};
const result = adjacencyListToMatrix(adjList);
assert.sameDeepOrderedMembers(result, [[0, 1], [0, 0]]);
```

Quando dada a lista de adjacência `{0: [1, 2], 1: [2], 2: [0, 3], 3: [2]}`, a função deve retornar `[[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]]`.

```js
const adjList = {0: [1, 2], 1: [2], 2: [0, 3], 3: [2]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1, 1, 0], [0, 0, 1, 0], [1, 0, 0, 1], [0, 0, 1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quando fornecida a lista de adjacência `{0: [1], 1: [0]}`, a função deve retornar `[[0, 1], [1, 0]]`.

```js
const adjList = {0: [1], 1: [0]};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 1], [1, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

Quando fornecida a lista de adjacência `{0: [], 1: [], 2: []}`, a função deve retornar `[[0, 0, 0], [0, 0, 0], [0, 0, 0]]`.

```js
const adjList = {0: [], 1: [], 2: []};
const result = adjacencyListToMatrix(adjList);
const expected = [[0, 0, 0], [0, 0, 0], [0, 0, 0]];
assert.sameDeepOrderedMembers(result, expected);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function adjacencyListToMatrix(adjList) {
  const n = Object.keys(adjList).length;

  const adjMatrix = Array.from({ length: n }, () => Array(n).fill(0));

  for (const [srcStr, neighbors] of Object.entries(adjList)) {
    const srcNode = Number(srcStr);
    for (const destNode of neighbors) {
      adjMatrix[srcNode][destNode] = 1;
    }
  }

  for (const row of adjMatrix) {
    console.log(row);
  }

  return adjMatrix;
}

const adjList = {
  0: [1, 2],
  1: [2],
  2: [0, 3],
  3: [2]
};

const matrix = adjacencyListToMatrix(adjList);
```
