---
id: 656485a8a3496d4a36b1496a
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

这三个变量构成了生成密码时可以选择的字符。

就在它们之前，添加一条注释 `Define the possible characters for the password`。

# --hints--

You should add the comment just above the `letters` variable.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
