---
id: 6564c67db26c417561ab510d
title: 步骤 31
challengeType: 20
dashedName: step-31
---

# --description--

从输出可以看到，现在你的正则匹配字符串中的第一个 `l`。

在你的正则表达式中，你可以在字符后添加一个量词，以指定该字符应重复多少次。 例如，`+` 量词意味着它应重复一次或更多次。

Add a `+` quantifier to your pattern.

# --hints--

You should modify your `pattern` variable into `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)

        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
