---
id: 6564883669b5af4b69f794cc
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

你的 `all_characters` 变量是一个由所有小写字母和大写字母、所有 10 个数字和几个特殊字符组成的字符串。

就在它之前，添加一条注释 `Combine all characters`。

# --hints--

你应该添加该注释，就在你的 `all_characters` 变量上方。

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
