---
id: 66a9cb26ec6bd41cf6c82bc5
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

La última sección a añadir a la tabla sería el elemento <dfn>table foot</dfn>, `tfoot`. El elemento de pie de tabla se utilizará para mostrar la calificación promedio de todos los estudiantes en la tabla.

Añadir un elemento `tfoot` a la tabla. 

Dentro del elemento `tfoot`, añadir un elemento `tr`.

Dentro del elemento `tr`, añadir dos elementos `td`. 

El primer elemento `td` debe contener el texto `Average Grade`. El segundo elemento `td` debe contener el texto `78.8`.

# --hints--

Debes tener una etiqueta de apertura `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Debes tener una etiqueta de cierre `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Debes tener un elemento `tr` dentro de tu elemento `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Tu primer elemento `td` debe contener el texto `Average Grade`. Asegúrate de que está dentro del elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Tu segundo elemento `td` debe contener el texto `78.8`. Asegúrate de que está dentro del elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Tu elemento `tr` debe tener dos elementos `td` dentro de él.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
