---
id: 66a9b48a211a73155621d0df
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este taller, practicarás trabajar con tablas HTML construyendo una tabla de examen final para un grupo de estudiantes. 

En lecciones anteriores, aprendiste cómo trabajar con el elemento `table` para representar datos tabulares.

Dentro de tu elemento `body`, anida un elemento `table`.

# --hints--

Debes tener una etiqueta de apertura `table`.

```js
assert.match(code, /<table>/i);
```

Debes tener una etiqueta de cierre `table`.

```js
assert.match(code, /<\/table>/i);
```

Tu elemento `table` debe estar dentro de tu elemento `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
