---
id: 66a9c14e3b34c719e34bae20
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Para agregar datos de estudiantes a la tabla, necesitarás usar los elementos de fila de tabla y de datos de tabla.

El elemento <dfn>table data</dfn>, `td`, se utiliza para crear una celda en la tabla.

Aquí tienes un ejemplo de uso del elemento `td` para una tabla de jugadores de deportes:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Dentro de tu elemento de cuerpo de tabla, agrega un elemento de fila de tabla, `tr`, con tres elementos de datos de tabla, `td`. 

El primer elemento de datos de tabla debe contener el apellido `Davis`. 

El segundo elemento de datos de tabla debe contener el nombre `Alex`. 

El tercer elemento de datos de tabla debe contener la calificación `54`.

# --hints--

Debes tener un elemento `tr` dentro de tu elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Debes tener un elemento `td` con el texto `Davis` dentro del elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Debes tener un elemento `td` con el texto `Alex` dentro del elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Debes tener un elemento `td` con el texto `54` dentro del elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Debes tener los tres elementos `td` dentro del elemento `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
