---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Ahora es el momento de agregar dos estudiantes más a la tabla.

Siguiendo el mismo patrón que el paso anterior, añade una segunda fila de tabla para estudiantes. Utiliza los siguientes datos para los elementos de datos de la tabla:

- Apellido: `Doe`
- Nombre: `Samantha`
- Calificación: `92`

Para la tercera fila de estudiantes de la tabla, utiliza los siguientes datos para los elementos de datos de la tabla:

- Apellido: `Rodriguez`
- Nombre: `Marcus`
- Calificación: `88`

# --hints--

Deberías tener un segundo `tr` dentro de tu elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Deberías tener un elemento `td` con el texto `Doe` dentro del segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Deberías tener un elemento `td` con el texto `Samantha` dentro del segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Deberías tener un elemento `td` con el texto `92` dentro del segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Deberías tener un tercer `tr` dentro de tu elemento `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Deberías tener un elemento `td` con el texto `Rodriguez` dentro del tercer elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Deberías tener un elemento `td` con el texto `Marcus` dentro del tercer elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Deberías tener un elemento `td` con el texto `88` dentro del tercer elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
