---
id: 66a9bb8578459a175432e7d0
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Para la primera sección de la tabla, querrás agrupar el contenido del encabezado que representa las etiquetas de columna para el nombre del estudiante, apellido y calificación en el examen final.

El elemento de <dfn>cabeza de tabla</dfn>, `thead`, se utiliza para agrupar el contenido del encabezado en una tabla. 

Aquí tienes un ejemplo usando el elemento `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Debajo de tu elemento `caption`, agrega un elemento para cabeza de tabla.

# --hints--

Debes tener una etiqueta de apertura `thead`.

```js
assert.match(code, /<thead>/i);
```

Debes tener una etiqueta de cierre `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Tu elemento `thead` debe estar dentro de tu elemento `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
