---
id: 66a9bcf00f13a418368a272e
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El elemento de la cabeza de la tabla consiste en un elemento <dfn>fila de tabla</dfn>, `tr`, que contiene los elementos <dfn>celda de encabezado de tabla</dfn>, `th`.

Aquí hay un ejemplo que utiliza los elementos `tr` y `th` para una tabla de deportes:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Dentro de tu elemento `thead`, añade un elemento `tr`.

Dentro de tu elemento `tr`, añade tres elementos `th`. 

El primer elemento `th` debe contener el texto `Last Name`. El segundo elemento `th` debe contener el texto `First Name`. El tercer elemento `th` debe contener el texto `Grade`.

# --hints--

Debes tener una etiqueta de apertura `tr`.

```js
assert.match(code, /<tr>/i);
```

Debes tener una etiqueta de cierre `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Tu elemento `tr` debe estar dentro de tu elemento `thead`. 

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Tu primer elemento `th` debe tener el texto de `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Tu segundo elemento `th` debe tener el texto de `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Tu tercer elemento `th` debe tener el texto de `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Todos tus elementos `th` deben estar dentro de tu elemento `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
