---
id: 5a24c314108439a4d403614a
title: Kusonga mbele kutoka hapa
challengeType: 6
forumTopicId: 301434
dashedName: moving-forward-from-here
---

# --description--

Hongera! Umehitimisha mafundisho kuhusu React na Redux. Kuna jambo moja la mwisho linalostahili kutajwa kabla hujaendelea. Kawaida, hutaunda programu za React katika mhariri wa msimbo kama huu. Changamoto hii inakupa mtazamo wa jinsi sintaksia inavyoonekana ikiwa unafanya kazi na mfumo wa faili kwenye kompyuta yako mwenyewe. Msimbo unapaswa kuonekana kama huu, isipokuwa kwa matumizi ya `import` (hizi huleta utegemezi wote uliotolewa kwako katika changamoto).

Mwishowe, kuandika msimbo wa React na Redux kwa kawaida kunahitaji usanidi fulani. Hii inaweza kuwa ngumu haraka. Ikiwa unavutiwa na kujaribu kwenye kompyuta yako mwenyewe, <a href="https://www.freecodecamp.org/news/install-react-with-create-react-app/" target="_blank" rel="noopener noreferrer nofollow">Create React App</a> inakuja tayari na imewekwa.

Vinginevyo, unaweza kuwezesha Babel kama Preprocessor wa JavaScript katika CodePen, ongeza React na ReactDOM kama rasilimali za JavaScript za nje, na ufanye kazi hapo pia.

# --instructions--

Andika ujumbe `'Now I know React and Redux!'` kwenye konsoli.

# --hints--

Ujumbe `Now I know React and Redux!` unapaswa kuandikwa kwenye konsoli.

```js
assert(
    /console\s*\.\s*log\s*\(\s*('|"|`)Now I know React and Redux!\1\s*\)/.test(
      code)
  );
```

# --seed--

## --seed-contents--

```jsx
/*
import React from 'react'
import ReactDOM from 'react-dom'
import { Provider, connect } from 'react-redux'
import { createStore, combineReducers, applyMiddleware } from 'redux'
import thunk from 'redux-thunk'

import rootReducer from './redux/reducers'
import App from './components/App'

const store = createStore(
  rootReducer,
  applyMiddleware(thunk)
);

ReactDOM.render(
  <Provider store={store}>
    <App/>
  </Provider>,
  document.getElementById('root')
);
*/

// Only change code below this line
```

# --solutions--

```jsx
console.log('Now I know React and Redux!');
```
