---
id: 5a24c314108439a4d4036143
title: Hamisha mantiki ya hali kwenda Redux
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

Sasa umeimaliza sehemu ya React, unahitaji kuhamisha mantiki inayotekelezwa ndani ya `state` kwenda Redux. Hii ni hatua ya kwanza kuunganisha programu rahisi ya React na Redux. Kazi pekee ya programu yako ni kuongeza ujumbe mpya kutoka kwa mtumizi kwenye orodha isiyopangwa. Mfano ni rahisi ili kuonyesha jinsi React na Redux vinavyofanya kazi pamoja.

# --instructions--

Kwanza, fafanua aina ya kitendo `ADD` na uiweke kuwa const `ADD`. Kisha, fafanua mtengenezaji wa kitendo `addMessage()` ambaye huunda kitendo cha kuongeza ujumbe. Utahitaji kupitisha `message` kwa mtengenezaji huyu wa kitendo na kujumuisha ujumbe katika `action` inayorejeshwa.

Kisha tengeneza reducer inayoitwa `messageReducer()` inayoshughulikia hali ya ujumbe. Hali ya awali inapaswa kuwa safu ya data tupu. Reducer hii inapaswa kuongeza ujumbe kwenye safu ya ujumbe iliyohifadhiwa katika hali, au kurejesha hali ya sasa. Mwishowe, tengeneza hifadhi yako ya Redux na uipitishe reducer.

# --hints--

Const `ADD` inapaswa kuwepo na kuwa na thamani sawa na mfuatano wa herufi `ADD`

```js
assert(ADD === 'ADD');
```

Mtengenezaji wa kitendo `addMessage` anapaswa kurejesha kitu chenye `type` sawa na `ADD` na `message` sawa na ujumbe unaopitishwa.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` inapaswa kuwa kitendakazi.

```js
assert(typeof messageReducer === 'function');
```

Hifadhi inapaswa kuwepo na kuwa na hali ya awali iliyowekwa kuwa safu ya data tupu.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

Kutuma `addMessage` dhidi ya hifadhi kunapaswa kuongeza ujumbe mpya kwa usalama kwenye safu ya ujumbe iliyohifadhiwa katika hali.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer` inapaswa kurejesha hali ya sasa ikiwa itaitwa na vitendo vingine vyovyote.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
