---
id: 68e76b1ce8216249d04df670
title: 建置一個 Planet 類別
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `Planet` 的類別。
1. `Planet` 類別應該有一個 `__init__` 方法，該方法：
   - 有四個參數：`self`、`name`、`planet_type` 和 `star`。
   - 如果傳入的任何引數不是字串型別，則會引發帶有訊息 `name, planet type, and star must be strings` 的 `TypeError`。
   - 如果傳入的任何引數是空字串，則會引發帶有訊息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。
   - 將傳入的值指定給實體屬性 `name`、`planet_type` 和 `star`。
1. `Planet` 類別應該有一個 `orbit` 方法，該方法傳回格式為 `{name} is orbiting around {star}...` 的字串。
1. `Planet` 類別應該有一個 `__str__` 方法，該方法傳回格式為 `Planet: {name} | Type: {planet_type} | Star: {star}` 的字串。
1. 你應該創建三個名為 `planet_1`、`planet_2` 和 `planet_3` 的 `Planet` 類別實體。
1. 你應該列印每個行星物件以查看 `__str__` 方法輸出。
1. 你應該對每個行星物件呼叫 `orbit` 方法並列印結果。

# --hints--

你應該創建一個名為 `Planet` 的類別。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

`Planet` 類別應該有一個 `__init__` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

`__init__` 方法應該有四個參數：`self`、`name`、`planet_type` 和 `star`，依此順序。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

如果傳入的第一個引數不是字串型別，`__init__` 方法應該引發帶有訊息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果傳入的第二個引數不是字串型別，`__init__` 方法應該引發帶有訊息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果傳入的第三個引數不是字串型別，`__init__` 方法應該引發帶有訊息 `name, planet type, and star must be strings` 的 `TypeError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果傳入的第一個引數是空字串，`__init__` 方法應該引發帶有訊息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果傳入的第二個引數是空字串，`__init__` 方法應該引發帶有訊息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

如果傳入的第三個引數是空字串，`__init__` 方法應該引發帶有訊息 `name, planet_type, and star must be non-empty strings` 的 `ValueError`。

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

你應該在 `__init__` 方法中將 `name` 指定到 `self.name`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

你應該在 `__init__` 方法中將 `planet_type` 指定給 `self.planet_type`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

你應該在 `__init__` 方法中將 `star` 指定到 `self.star`。

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

`Planet` 類別應該有一個 `orbit` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

`orbit` 方法應該傳回格式為 `{name} is orbiting around {star}...` 的字串。

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

`Planet` 類別應該有一個 `__str__` 方法。

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

`__str__` 方法應該傳回格式為 `Planet: {name} | Type: {planet_type} | Star: {star}` 的字串。

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

你應該創建三個名為 `planet_1`、`planet_2` 和 `planet_3` 的 `Planet` 類別實體。

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

你應該列印每個行星物件以查看 `__str__` 方法輸出。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

你應該對每個行星物件呼叫 `orbit` 方法並列印結果。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
