---
id: 587d78af367417b2b2512b03
title: 建置調查表單
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**目標：** 建置一個功能上類似於 <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a> 的應用程式。**請勿複製此示範專案**。

**使用者故事：**

1. 你應該在一個具有 `id` 為 `title` 的 `h1` 元素中放置頁面標題。
1. 你應該在具有 `id` 為 `description` 的 `p` 元素中放置簡短說明。
1. 你應該有一個 `form` 元素，其 `id` 為 `survey-form`
1. 在表單元素內，你**必須**在具有 `id` 為 `name` 且 `type` 為 `text` 的 `input` 欄位中輸入你的名字。
1. 在表單元素內，你**必須**在具有 `id` 為 `email` 的 `input` 欄位中輸入你的電子郵件。
1. 如果你輸入的電子郵件格式不正確，你將會看到 HTML5 驗證錯誤。
1. 在表單內，你可以在具有 `id` 為 `number` 的 `input` 欄位中輸入數字。
1. 數字輸入不應接受非數字，無論是透過阻止你輸入，或是顯示 HTML5 驗證錯誤（視你的瀏覽器而定）。
1. 如果你輸入的數字超出由 `min` 和 `max` 屬性定義的數字輸入範圍，將會看到 HTML5 驗證錯誤。
1. 對於名稱、電子郵件和數字輸入欄，你可以在表單中看到對應的 `label` 元素，這些元素使用以下 id 描述每個欄的用途：`id="name-label"`、`id="email-label"` 和 `id="number-label"`
1. 對於名稱、電子郵件和號碼輸入欄，你可以看到提供每個欄描述或指示的佔位文字。
1. 在表單元素內，你應該有一個 `select` 下拉選單元素，其 `id` 為 `dropdown`，並且至少有兩個可供選擇的選項。
1. 在表單元素內，你可以從使用 `name` 屬性分組的至少兩個圓鈕群組中選擇一個選項。
1. 在表單元素內，你可以從一系列核取方塊中選擇多個欄，每個核取方塊都必須有一個 `value` 屬性。
1. 在表單元素內，你會看到一個用於額外評論的 `textarea`。
1. 在表單元素內，你會看到一個 `id` 為 `submit` 的按鈕，用來提交所有輸入。

完成使用者故事並通過以下所有測試以完成此專案。展現你個人的風格。祝你程式設計愉快！

**注意：** 請務必在你的 HTML 中添加 `<link rel="stylesheet" href="styles.css">` 以聯結你的樣式表並套用你的 CSS。

# --hints--

你應該有一個 `h1` 元素，其 `id` 為 `title`。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

你的 `#title` 不應該是空的。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你應該有一個 `p` 元素，其 `id` 為 `description`。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

你的 `#description` 不應該是空的。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你應該有一個 `form` 元素，其 `id` 為 `survey-form`。

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

你應該有一個 `input` 元素，其 `id` 為 `name`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#name` 應該有 `text` 的 `type`。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

你的 `#name` 應該需要輸入。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#name` 應該是 `#survey-form` 的後代。

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

你應該有一個 `input` 元素，其 `id` 為 `email`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#email` 應該有 `email` 的 `type`。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

你的 `#email` 應該要求輸入。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

你的 `#email` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

你應該有一個 `input` 元素，其 `id` 為 `number`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

你的 `#number` 應該是 `#survey-form` 的後代。

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

你的 `#number` 應該有 `number` 的 `type`。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

你的 `#number` 應該有一個帶有數值的 `min` 屬性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

你的 `#number` 應該有一個帶有數值的 `max` 屬性。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

你應該有一個 `label` 元素，其 `id` 為 `name-label`。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你應該有一個 `label` 元素，其 `id` 為 `email-label`。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你應該有一個 `label` 元素，其 `id` 為 `number-label`。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

你的 `#name-label` 應該包含描述輸入的文字。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#email-label` 應包含描述輸入的文字。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#number-label` 應包含描述輸入的文字。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

你的 `#name-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

你的 `#email-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

你的 `#number-label` 應該是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

你的 `#name` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#email` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你的 `#number` 應該有一個 `placeholder` 屬性和值。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

你應該有一個 `select` 欄位，其 `id` 為 `dropdown`。

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

你的 `#dropdown` 應該至少有兩個可選擇（未停用）的 `option` 元素。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

你的 `#dropdown` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

你應該至少有兩個 `input` 元素，其 `type` 為 `radio`（圓鈕、單選按鈕）。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你應該至少有兩個作為 `#survey-form` 子孫的圓鈕。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

你所有的圓鈕都應該有一個 `value` 屬性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你所有的圓鈕都應該有一個 `name` 屬性和值。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

每個圓鈕群組應該至少有 2 個圓鈕。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

你應該至少有兩個 `input` 元素，其 `type` 為 `checkbox`（核取方塊），且是 `#survey-form` 的子孫元素。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 內所有的核取方塊都應該有 `value` 屬性和值。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

你應該至少有一個 `textarea` 元素是 `#survey-form` 的子孫。

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

你應該有一個 `id` 為 `submit` 的 `input` 或 `button` 元素。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

你的 `#submit` 應該有 `submit` 的 `type`。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

你的 `#submit` 應該是 `#survey-form` 的子孫元素。

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
