---
id: 6823f9df49cc206af5471a30
title: Rekebisha ukurasa wa wasifu wa Camperbot
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Camperbot anajaribu kutengeneza ukurasa wa wasifu. Aliomba rafiki yake akague msimbo wake na rafiki yake alisema kuna makosa kadhaa.

Kazi yako ni kurekebisha makosa yote ya Camperbot ili aweze kuendelea kutengeneza ukurasa wake wa wasifu. Kamilisha vitu vilivyo kwenye hadithi za mtumizi hapa chini na bofya "Check Your Code" kuona kama umeondoa makosa yote.

**Hadithi za mtumizi:**

1. Camperbot anajaribu kutumia kipengele cha `heading2`, lakini kipengele hicho hakipo. Rekebisha lebo hizo ili zitumie kichwa cha sehemu cha ngazi ya pili kinachofaa.
2. Camperbot anajaribu kuongeza aya mbili kwa kutumia `pp`, lakini hizo pia hazipo. Rekebisha ili zitumie lebo za aya zinazofaa.
3. Camperbot anatumia kipengele cha `h3` kwa kichwa kidogo cha `Background and Interests` lakini kina kosa la sintaksia. Tambua tatizo na ulitatue.

# --hints--

Unapaswa kuwa na kipengele kimoja tu cha `h1` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Haupaswi kuondoa kipengele cha `<h1>Hello from Camperbot!</h1>` kutoka kwenye ukurasa.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

Unapaswa kuwa na kipengele kimoja tu cha `h2` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Unapaswa kuwa na kipengele cha `h2` chenye maandishi ya `About`. Hapa kuna mfano: `<h2>Text here</h2>`

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Unapaswa kuwa na kipengele cha aya chenye maandishi ya `My name is Camperbot and I love learning new things.` Hapa kuna mfano: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Unapaswa kuwa na kipengele cha aya chenye maandishi ya `I enjoy solving puzzles.` Hapa kuna mfano: `<p>Text here</p>`

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

Unapaswa kuwa na vipengele viwili vya aya kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

Unapaswa kuwa na kipengele kimoja tu cha `h3` kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Unapaswa kuwa na kipengele cha `h3` chenye maandishi ya `Background and Interests`. Hapa kuna mfano: `<h3>Text here</h3>`

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Unapaswa kuwa na vipengele vitano jumla kwenye ukurasa.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
