---
id: 66bba6fff611169359d9d36a
title: 秘密のメールページを作成する
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**目的：** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `main` 要素に `id` が `email` であること。
1. `#email` 要素は `padding` が `50px`、上マージンが `50px`、`width` が `500px`、そして `border` が `2px` であること。
1. `#email` 要素の合計幅（パディングとボーダーを含む）が `500px` であること。
1. `#email` 要素内に、id が `confidential` の `div` 要素と、id が `top-secret` の `div` 要素が2つあること。
1. `#confidential` と `#top-secret` 要素は `display` が `inline-block` であること。
1. `#confidential` と `#top-secret` 要素は `padding`、左マージン、そして `border` を持つこと。
1. `#confidential` 要素のテキストが `CONFIDENTIAL` であること。
1. `#top-secret` 要素のテキストが `TOP SECRET` であること。
1. `#confidential` と `#top-secret` 要素は CSS の transform を使って回転していること。
1. `#email` 要素内に少なくとも3つの段落要素があること。
1. 段落要素内にクラスが `blurred` の `span` 要素が少なくとも3つあること。
1. `blurred` セレクターが CSS フィルターを使って `3px` 要素をぼかしていること。

# --hints--

`main` 要素に `id` が `email` であること。

```js
assert.exists(document.querySelector('main#email'));
```

`#email` セレクターが、その要素のすべての辺の `padding` を `50px` に設定していること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

`#email` セレクターが、その要素の `margin-top` を `50px` に設定していること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

`#email` セレクターが、その要素の `width` を `500px` に設定していること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

`#email` 要素に `2px` のボーダーがあること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

`#email` 要素に `box-sizing` が `border-box` であること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

`div` 要素に `id` が `confidential` のものが `#email` 要素内にあること。

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

`div` 要素に `id` が `top-secret` のものが `#email` 要素内にあること。

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

`#confidential` 要素の `display` が `inline-block` に設定されていること。

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

`#top-secret` 要素の `display` が `inline-block` に設定されていること。

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

`#confidential` 要素にすべての辺に `padding` があること。

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

`#top-secret` 要素にすべての辺に `padding` があること。

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

`#confidential` 要素に左マージンがあること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

`#top-secret` 要素に左マージンがあること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

`#confidential` 要素に `border` があること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

`#top-secret` 要素に `border` があること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

`#confidential` 要素のテキストが `CONFIDENTIAL` であること。テキストが抜けているか、タイプミスがあります。

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

`#top-secret` 要素のテキストが `TOP SECRET` であること。テキストが抜けているか、タイプミスがあります。

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

`#confidential` 要素は CSS の `transform` を使って回転していること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

`#top-secret` 要素は CSS の `transform` を使って回転していること。

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

`#email` 要素内に少なくとも3つの段落要素があること。

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

段落内にクラスが `blurred` の `span` 要素が少なくとも3つあること。

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

`.blurred` 要素は空であってはいけません。

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

`.blurred` セレクターが、その要素の `filter` を `blur(3px)` に設定していること。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
