---
id: 66b102ee0bdbad34a9f85ba0
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、Teacher Chatbotを作りながら文字列についてさらに学んでいきます。

まず、メッセージに`console`を使った`"Hi there!"`文を追加してください。

# --hints--

`console`文があるはずです。

```js
assert.match(code, /console\.log(.*)/);
```

コンソールの文はメッセージ`"Hi there!"`を出力する必要があります。

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
