---
id: 66b59829dba144ff1351220f
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`botName` 変数を使って挨拶文を作成する時間です。

これまでのレッスンでは、テンプレートリテラルを使って文字列を連結する方法を学びました。例えば次のように：

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

まず、`greeting` という名前の変数を作成してください。

次に、テンプレートリテラルの構文を使って、「`My name is`」という文字列に `botName` 変数を続け、最後にピリオド（`.`）を付けた文字列を代入してください。

最後に、`greeting` 変数をコンソールに出力してください。

# --hints--

`greeting` という名前の変数があるはずです。

```js
assert.isNotNull(greeting);
```

`greeting` 変数に文字列を代入してください。

```js
assert.isString(greeting);
```

テンプレートリテラルを使って、文字列 `My name is` と `botName` 変数、そしてピリオド（`.`）を連結してください。

```js
assert.equal(greeting, "My name is teacherBot.");
```

`console` 変数をコンソールに出力する `greeting` 文があるはずです。

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
