---
id: 66b59b12a745e10011158f7b
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

次のステップは、将来のボットメッセージで使ういくつかの変数を作成することです。

`subject` という変数を作成し、文字列の値 `"JavaScript"` を代入してください。

次に、`topic` という変数を作成し、文字列の値 `"strings"` を代入してください。

# --hints--

`subject` という変数があるはずです。

```js
assert.isNotNull(subject);
```

`subject` 変数には文字列を代入してください。

```js
assert.isString(subject);
```

`"JavaScript"` 変数に文字列 `subject` を代入してください。

```js
assert.equal(subject, "JavaScript");
```

`topic` という変数があるはずです。

```js
assert.isNotNull(topic);
```

`topic` 変数には文字列を代入してください。

```js
assert.isString(topic);
```

`"strings"` 変数に文字列 `topic` を代入してください。

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
