---
id: 66b59be6ab830800c4df9146
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

前のステップで作成した変数を使う時が来ました。

これまでのレッスンでは、次のようなテンプレートリテラルの扱い方を学びました。

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

まず、`sentence` 変数を作成してください。

テンプレートリテラルの構文を使って、`Today, you will learn about [topic variable goes here] in [subject variable goes here].` という文字列を `sentence` 変数に代入してください。

`[topic variable goes here]` と `[subject variable goes here]` のプレースホルダーは、`topic` と `subject` の変数および `${}` の構文で置き換えます。

最後に、`sentence` 変数をコンソールに出力してください。

# --hints--

`sentence` 変数があるはずです。

```js
assert.isNotNull(sentence);
```

`sentence` 変数は文字列であるべきです。

```js
assert.isString(sentence);
```

テンプレートリテラルの構文を使って、`Today, you will learn about [topic variable goes here] in [subject variable goes here].` という文字列を `sentence` 変数に代入してください。助けが必要なら例を参照してください。

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

`sentence` 変数をコンソールに出力してください。

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
