---
id: 66b5928b86e907fdfbf34e56
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

ボットの名前を設定する時間です。

`botName` という変数を作成し、文字列の値 `"teacherBot"` を代入してください。

# --hints--

`botName` という変数があるはずです。

```js
assert.isNotNull(botName);
```

`botName` 変数は文字列である必要があります。

```js
assert.isString(botName);
```

`botName` 変数には `"teacherBot"` の値が入っている必要があります。

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
