---
id: 66b6e80d6c3f0b329c360283
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

このワークショップの次の部分では、文字列から文字をアクセスする方法をレビューします。

まず、メッセージ `Here is an example of accessing the first letter in the word [subject].` をコンソールに出力してください。

前のステップで行ったように、`[subject]` を `subject` 変数に置き換え、適切なテンプレートリテラル構文を使うことを忘れないでください。

# --hints--

あなたの `console` 文はメッセージ `Here is an example of accessing the first letter in the word [subject].` を出力する必要があります

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+first\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

--fcc-editable-region--

--fcc-editable-region--
```
