---
id: 66b6e62423e8a031d6c1f03d
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

今度は `topic` 文字列の長さを取得する時です。

`console` 文の中でテンプレートリテラルを次のように使うことができます。

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

まず、メッセージ `Here is an example of using the length property on the word [topic].` をコンソールに出力してください。

`[topic]` を `topic` 変数に置き換え、前のステップで行ったように正しいテンプレートリテラルの構文を使うことを忘れないでください。

次に、2つ目の `console.log` 文を追加して、`topic` 文字列の長さをコンソールに出力してください。

# --hints--

メッセージ `Here is an example of using the length property on the word [topic].` を `console` に出力するにはテンプレートリテラルの構文を使うべきです。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

2つ目の `console` 文で `topic` 変数の長さを `console` に出力しているはずです。

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
