---
id: 66b6e39031393e30f2c48d0f
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

文字列の長さを取得するには、<dfn>length</dfn> プロパティを使います。このプロパティは文字列の文字数を返します。

例を示します：

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

まず、`length` プロパティを使って `subject` 文字列の長さを取得し、その値をコンソールに出力してください。

# --hints--


`length` 文字列の長さを取得するには `subject` プロパティを使う必要があります。

```js
assert.match(code, /subject\.length/);
```

`subject` 文字列の長さをコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
