---
id: 66b6ffb42b88e33943788abf
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

ワークショップの最後の部分では、文字列の中の部分文字列のインデックス位置を見つける方法をレビューします。部分文字列とは文字列の一部のことです。

まず、`learningIsFunSentence` という変数を作成し、文字列の値として `"Learning is fun."` を代入してください。

# --hints--

`learningIsFunSentence` という変数を用意してください。

```js
assert.isNotNull(learningIsFunSentence);
```

`learningIsFunSentence` には文字列の値が入っている必要があります。

```js
assert.isString(learningIsFunSentence);
```

`"Learning is fun."` の値を変数 `learningIsFunSentence` に代入してください。

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
