---
id: 66b6d482bbb9e12f2e5ee1ae
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

このプロジェクトの次の部分では、ボットが文字列 `length` プロパティの扱い方を教えます。

まず、`strLengthIntro` という新しい変数を作成してください。

次に、テンプレートリテラル構文を使って、文字列 `Here is an example of using the length property on the word [subject].` を `strLengthIntro` 変数に代入してください。

先ほどと同じように、`[subject]` を `subject` 変数に置き換えてください。

最後に、`strLengthIntro` 変数をコンソールにログ出力してください。

# --hints--

`strLengthIntro` という変数があるはずです。

```js
assert.isNotNull(strLengthIntro);
```

`strLengthIntro` 変数は文字列であるべきです。

```js
assert.isString(strLengthIntro);
```

テンプレートリテラル構文を使って、文字列 `Here is an example of using the length property on the word [subject goes here].` を `strLengthIntro` 変数に代入してください。

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

`strLengthIntro` 変数をコンソールにログ出力してください。

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
