---
id: 66b6efddeca35833cd6f0b03
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

前のレッスンでは、文字列の文字に次のようにアクセスする方法を学びました。

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

インデックス番号は`0`から始まるので、文字列の最初の文字は常にインデックス`0`にあります。

まず、もう一つ`console`文を追加してください。

`console`文の中で、ブラケット記法と正しいインデックス番号を使って、`subject`変数の最初の文字を出力してください。

# --hints--

`console`文は、ブラケット記法とインデックス`subject`を使って、`0`変数の最初の文字を出力する必要があります。

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
