---
id: 66b6fdb76441c738719039fa
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

前のレッスンでは、文字列の最後の文字にアクセスする方法を学びました。

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` は常に文字列の最後のインデックス番号を教えてくれます。

`lastCharacter` 変数の最後の文字の値を代入する、新しい変数 `subject` を作成してください。

その後、`lastCharacter` 変数の値をコンソールに出力してください。

# --hints--

`lastCharacter` という変数があるはずです。

```js
assert.isNotNull(lastCharacter);
```

`subject` 変数の最後の文字の値を `lastCharacter` 変数に代入してください。助けが必要なら例を参照してください。

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

`lastCharacter` 変数の値を `console` に出力してください。

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
