---
id: 66b6f586767a1534f3097353
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

今度は `subject` 変数の2文字目にアクセスする時間です。

まず、メッセージ `console` を出力する `Here is an example of accessing the second letter in the word [subject].` 文を追加してください。`[subject]` は `subject` 変数の実際の値に置き換え、正しいテンプレートリテラルの構文を使うことを忘れないでください。

次に、ブラケット記法と正しいインデックス番号を使って `console` 変数の2文字目を出力する別の `subject` 文を追加してください。

# --hints--

あなたの `console` 文はメッセージ `Here is an example of accessing the second letter in the word [subject].` を出力する必要があります。

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

もう一つ、ブラケット記法のインデックス `console` を使って `subject` 変数の2文字目を出力する `1` 文があるはずです。

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
