---
id: 66b7137348cfb53fd3ec6c73
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

次に、`console`変数に対して`indexOf`メソッドを使い、部分文字列`learningIsFunSentence`の位置を見つける結果を出力する新しい`"fun"`文を追加してください。

その`console`文の下に、部分文字列`console`の位置を見つけるために`indexOf`メソッドを使った結果を出力する新しい`"learning"`文を追加してください。

最後の`console`文が出力する内容に注目してください。

# --hints--

あなたの`console`は、`"fun"`変数から部分文字列`learningIsFunSentence`のインデックス位置を出力する必要があります。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

2つ目の`console`文は、`"learning"`変数から部分文字列`learningIsFunSentence`のインデックス位置を出力する必要があります。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
