---
id: 66b70b8a611cbf3bcc5c6c5f
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

前のレッスンでは、次のように`indexOf`メソッドを使う方法を学びました。

```js
const sentence = "I love to learn.";

// returns index 2
console.log(sentence.indexOf("love"));

// returns -1
console.log(sentence.indexOf("hate"));
```

`indexOf`メソッドは、文字列の中で部分文字列が最初に現れるインデックス位置を返すことを覚えておいてください。部分文字列が見つからない場合は、`-1`を返します。

`console`文を追加して、`indexOf`変数に対して`learningIsFunSentence`メソッドを使い、部分文字列`"Learning"`のインデックス位置を取得した結果を出力してください。

# --hints--

`console`は、`indexOf`変数に対して`learningIsFunSentence`メソッドを使い、部分文字列`"Learning"`のインデックス位置を見つけるようにしてください。

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(Learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

--fcc-editable-region--

--fcc-editable-region--
```
