---
id: 579e2a2c335b9d72dd32e05c
title: スライスとスプライス
challengeType: 1
forumTopicId: 301148
dashedName: slice-and-splice
---

# --description--

2つの配列とインデックスが与えられます。

最初の配列の各要素を順番に2番目の配列にコピーしてください。

2番目の配列のインデックス`n`から要素の挿入を始めてください。

結果の配列を返してください。関数実行後も入力の配列は変わらないようにしてください。

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)`は`[4, 1, 2, 3, 5]`を返すべきです。

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)`は`["a", 1, 2, "b"]`を返すべきです。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)`は`["head", "shoulders", "claw", "tentacle", "knees", "toes"]`を返すべきです。

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

最初の配列のすべての要素は元の順序のまま2番目の配列に追加されるべきです。`frankenSplice([1, 2, 3, 4], [], 0)`は`[1, 2, 3, 4]`を返すべきです。

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

関数実行後、最初の配列は変わらないようにしてください。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

関数実行後、2番目の配列は変わらないようにしてください。

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js
function frankenSplice(arr1, arr2, n) {
  return arr2;
}

frankenSplice([1, 2, 3], [4, 5, 6], 1);
```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  // It's alive. It's alive!
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}

frankenSplice([1, 2, 3], [4, 5], 1);
```
