---
id: a202eed8fc186c8434cb6d61
title: 文字列を逆にする
challengeType: 1
forumTopicId: 16043
dashedName: reverse-a-string
---

# --description--

与えられた文字列を逆にして、逆にした文字列を返してください。

例えば、`"hello"` は `"olleh"` になります。

# --hints--

`reverseString("hello")` は文字列を返す必要があります。

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` は文字列 `olleh` を返す必要があります。

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` は文字列 `ydwoH` を返す必要があります。

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` は文字列 `htraE morf sgniteerG` を返す必要があります。

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js
function reverseString(str) {
  return str;
}

reverseString('hello');
```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
