---
id: a24c1a4622e3c05097f71d67
title: どこに属するか
challengeType: 1
forumTopicId: 16094
dashedName: where-do-i-belong
---

# --description--

値（第2引数）をソート済みの配列（第1引数）に挿入する場合、挿入すべき最も低いインデックスを返してください。返り値は数値である必要があります。

例えば、`getIndexToIns([1,2,3,4], 1.5)` は `1`（インデックス0）より大きく、`1`（インデックス1）より小さいため、`2` を返すべきです。

同様に、`getIndexToIns([20,3,5], 19)` は配列がソートされた後に `2` のようになり、`[3,5,20]` は `19`（インデックス2）より小さく、`20`（インデックス1）より大きいため、`5` を返すべきです。

# --hints--

`getIndexToIns([10, 20, 30, 40, 50], 35)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` は `3` を返すべきです。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 30));
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([40, 60], 50));
```

`getIndexToIns([40, 60], 50)` は `1` を返すべきです。

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([3, 10, 5], 3));
```

`getIndexToIns([3, 10, 5], 3)` は `0` を返すべきです。

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([5, 3, 20, 3], 5));
```

`getIndexToIns([5, 3, 20, 3], 5)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([2, 20, 10], 19)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([2, 20, 10], 19));
```

`getIndexToIns([2, 5, 10], 15)` は `3` を返すべきです。

```js
assert.strictEqual(getIndexToIns([2, 5, 10], 15), 3);
```

`getIndexToIns([2, 5, 10], 15)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([2, 5, 10], 15));
```

`getIndexToIns([], 1)` は数値を返すべきです。

```js
assert.isNumber(getIndexToIns([], 1));
```

`getIndexToIns([], 1)` は `0` を返すべきです。

```js
assert.strictEqual(getIndexToIns([], 1), 0);
```

# --seed--

## --seed-contents--

```js
function getIndexToIns(arr, num) {
  return num;
}

getIndexToIns([40, 60], 50);
```

# --solutions--

```js
function getIndexToIns(arr, num) {
  arr = arr.sort((a, b) => a - b);

  for (let i = 0; i < arr.length; i++) {
    if (arr[i] >= num) {
      return i;
    }
  }

  return arr.length;
}

getIndexToIns([40, 60], 50);
```
