---
id: a302f7aae1aa3152a5b413bc
title: 数の階乗を求める
challengeType: 1
forumTopicId: 16013
dashedName: factorialize-a-number
---

# --description--

与えられた整数の階乗を返してください。

整数が `n` で表される場合、階乗とは `n` 以下のすべての正の整数の積です。

階乗はしばしば省略記法 `n!` で表されます。

例えば：`5! = 1 * 2 * 3 * 4 * 5 = 120`

関数には 0 以上の整数のみが渡されます。

# --hints--

`factorialize(5)` は数値を返す必要があります。

```js
assert.isNumber(factorialize(5));
```

`factorialize(5)` は `120` を返す必要があります。

```js
assert.strictEqual(factorialize(5), 120);
```

`factorialize(10)` は `3628800` を返す必要があります。

```js
assert.strictEqual(factorialize(10), 3628800);
```

`factorialize(20)` は `2432902008176640000` を返す必要があります。

```js
assert.strictEqual(factorialize(20), 2432902008176640000);
```

`factorialize(0)` は `1` を返す必要があります。

```js
assert.strictEqual(factorialize(0), 1);
```

# --seed--

## --seed-contents--

```js
function factorialize(num) {
  return num;
}

factorialize(5);
```

# --solutions--

```js
function factorialize(num) {
  return num < 1 ? 1 : num * factorialize(num - 1);
}

factorialize(5);
```
