---
id: a6e40f1041b06c996f7b2406
title: 見つけたものは自分のもの
challengeType: 1
forumTopicId: 16016
dashedName: finders-keepers
---

# --description--

配列 `arr` を調べて、「真偽テスト」に合格する最初の要素を返す関数を作成してください。これは、要素 `x` が与えられたとき、`func(x)` が `true` であれば「真偽テスト」に合格することを意味します。もしどの要素もテストに合格しなければ、`undefined` を返してください。

# --hints--

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` は `8` を返すべきです。

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` は `undefined` を返すべきです。

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js
function findElement(arr, func) {
  let num = 0;
  return num;
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
