---
id: a77dbc43c33f39daa4429b4f
title: ブール値かどうかを判定する
challengeType: 1
forumTopicId: 16000
dashedName: boo-who
---

# --description--

値がブール値プリミティブとして分類されるかどうかを確認してください。`true` または `false` を返してください。

ブール値プリミティブは `true` と `false` です。

# --hints--

`booWho(true)` は `true` を返すべきです。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` は `true` を返すべきです。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` は `false` を返すべきです。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` は `false` を返すべきです。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` は `false` を返すべきです。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` は `false` を返すべきです。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` は `false` を返すべきです。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` は `false` を返すべきです。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` は `false` を返すべきです。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` は `false` を返すべきです。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js
function booWho(bool) {
  return bool;
}

booWho(null);
```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

booWho(null);
```
