---
id: ab6137d4e35944e21037b769
title: 文の各単語をタイトルケースにする
challengeType: 1
forumTopicId: 16088
dashedName: title-case-a-sentence
---

# --description--

与えられた文字列の各単語の最初の文字を大文字にして返してください。残りの文字はすべて小文字にしてください。

この演習の目的上、`the`や`of`のような接続詞も大文字にしてください。

# --hints--

`titleCase("I'm a little tea pot")`は文字列を返す必要があります。

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")`は文字列`I'm A Little Tea Pot`を返す必要があります。

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")`は文字列`Short And Stout`を返す必要があります。

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")`は文字列`Here Is My Handle Here Is My Spout`を返す必要があります。

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js
function titleCase(str) {
  return str;
}

titleCase("I'm a little tea pot");
```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}

titleCase("I'm a little tea pot");
```
