---
id: 69a92f0b9cc04bb0d5327bb5
title: Implementar una pila
challengeType: 26
dashedName: implement-a-stack
---

# --description--

En este laboratorio, implementarás una estructura de datos pila usando funciones. Una pila es una estructura de datos LIFO (último en entrar, primero en salir) donde los elementos se agregan y se eliminan desde la parte superior.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función `initStack` que devuelva un objeto con una propiedad `collection` establecida en un arreglo vacío. El último elemento de `collection` representa la cima de la pila.
1. Debes tener una función `push` que tome un objeto pila y un elemento como argumentos, y agregue el elemento a la cima de la pila colocándolo al final de la `collection` de la pila.
1. Debes tener una función `pop` que tome un objeto pila como argumento, y elimine y devuelva el elemento superior de la pila, o `undefined` si no hay ninguno.
1. Debes tener una función `peek` que tome un objeto pila como argumento, y devuelva el elemento superior de la pila sin eliminarlo, o `undefined` si no hay ninguno.
1. Debes tener una función `isEmpty` que tome un objeto pila como argumento, y devuelva `true` si la pila no contiene elementos, y `false` en caso contrario.
1. Debes tener una función `clear` que tome un objeto pila como argumento, y elimine todos los elementos de la `collection` de la pila.

**Nota:** La mayoría de las pruebas dependen de las funciones `initStack` y `push`. Implementa primero estas funciones, ya que las pruebas para `pop`, `peek`, `isEmpty` y `clear` requieren agregar elementos a la pila.

# --hints--

Debes tener una función `push`.

```js
assert.isFunction(push);
```

La función `push` debe agregar un elemento a la cima de la pila.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Debes tener una función `pop`.

```js
assert.isFunction(pop);
```

La función `pop` debe eliminar y devolver el elemento superior de la pila.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.deepEqual(stack.collection, ['first', 'second']);
assert.strictEqual(pop(stack), 'second');
assert.deepEqual(stack.collection, ['first']);
```

La función `pop` debe devolver `undefined` si la pila está vacía.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

La función `pop` debe devolver correctamente los valores falsos.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Debes tener una función `peek`.

```js
assert.isFunction(peek);
```

La función `peek` debe devolver el elemento superior de la pila sin eliminarlo.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

La función `peek` debe devolver `undefined` si la pila está vacía.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

La función `peek` debe devolver correctamente los valores falsos.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Debes tener una función `isEmpty`.

```js
assert.isFunction(isEmpty);
```

La función `isEmpty` debe devolver `true` para una pila vacía.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

La función `isEmpty` debe devolver `false` para una pila no vacía.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

La función `isEmpty` debe devolver `false` cuando el elemento superior es un valor falso.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Debes tener una función `clear`.

```js
assert.isFunction(clear);
```

La función `clear` debe eliminar todos los elementos de la pila.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
