---
id: 67116c175a77cb64c5c10d49
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

Nelle lezioni precedenti, hai imparato a lavorare con funzioni di ordine superiore come `map`, `filter` e `reduce`.

In questo workshop, costruirai un'app per la gestione di una biblioteca che ti darà l'opportunità di esercitarti con diverse funzioni di ordine superiore.

Per iniziare, crea una variabile chiamata `library` e assegnale un array vuoto.

# --hints--

Dovresti creare un array `library`.

```js
assert.isArray(library)
```

Il tuo array `library` dovrebbe essere vuoto.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
