---
id: 67116d7584d0b469b14579bf
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

All'interno dell'array `library`, crea un oggetto con le seguenti proprietà e valori:

| Proprietà | Valore   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

Dovresti avere un oggetto all'interno dell'array `library`.

```js
assert.isObject(library[0])
```

Il tuo oggetto dovrebbe avere una proprietà `title`.

```js
assert.property(library[0], "title")
```

Il valore della tua proprietà `title` dovrebbe essere `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

Il tuo oggetto dovrebbe avere una proprietà `author`.

```js
assert.property(library[0], "author")
```

La proprietà `author` dovrebbe avere il valore `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

Il tuo oggetto dovrebbe avere una proprietà `about`.

```js
assert.property(library[0], "about")
```

La proprietà `about` dovrebbe avere il valore `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

Il tuo oggetto dovrebbe avere una proprietà `pages`.

```js
assert.property(library[0], "pages")
```

Il valore della proprietà `pages` dovrebbe essere un numero.

```js
assert.isNumber(library[0]?.pages)
```

La proprietà `pages` dovrebbe essere impostata su `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
