---
id: 67116d7584d0b469b14579c0
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Crea un altro oggetto all'interno dell'array `library` con le seguenti proprietà e valori:

| Proprietà | Valore   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


Dovresti avere un oggetto come secondo elemento del tuo array `library`.

```js
assert.isObject(library[1]);
```

Il tuo oggetto dovrebbe avere una proprietà `title`.

```js
assert.property(library[1], "title")
```

La proprietà `title` dovrebbe avere `"Atomic Habits"` come valore.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

Il tuo oggetto dovrebbe avere una proprietà `author`.

```js
assert.property(library[1], "author")
```

La proprietà `author` dovrebbe avere `"James Clear"` come valore.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

Il tuo oggetto dovrebbe avere una proprietà `about`.

```js
assert.property(library[1], "about")
```

La proprietà `about` dovrebbe avere il valore `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

Il tuo oggetto dovrebbe avere una proprietà `pages`.

```js
assert.property(library[1], "pages")
```

Il valore della proprietà `pages` dovrebbe essere un numero.

```js
assert.isNumber(library[1]?.pages)
```

La proprietà `pages` dovrebbe essere impostata su `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
