---
id: 697fc395f3d85f14b05cd9ca
title: استعادة سرد متماسك من مصفوفة من أجزاء القصة
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

في هذا المختبر، ستستعيد سردًا متماسكًا من مصفوفة تالفة من أجزاء القصة.

ستتمرن على استخدام الحلقات من خلال تنفيذ خوارزميات مصفوفة أساسية من الصفر.

ستعمل مع مصفوفات من كائنات أجزاء القصة. كل كائن جزء يحتوي على الخصائص التالية:

| الخاصية | الوصف | قيمة المثال |
|-|-|-|
| `id` | عدد صحيح موجب يحدد موضع الجزء في القصة | `3` |
| `text` | محتوى القصة الفعلي | `"and I use Arch btw.\""` |

في هذا المختبر، يُزوّد لك مصفوفة مملوءة مسبقًا تسمى `shuffledFragments`.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم**

1. لا يجب عليك تغيير مصفوفة `shuffledFragments` المملوءة مسبقًا.
2. يجب أن تنشئ دالة تسمى `compactFragments` تأخذ مصفوفة من الأجزاء وتعيد مصفوفة جديدة مع إزالة جميع عناصر undefined. إذا أزالت الدالة أي عناصر undefined، فيجب أن تسجل رسالة في وحدة التحكم. يجب أن تبدأ الرسالة بالبادئة `[COMPACTED]`.
3. يجب أن تعلن متغيرًا يسمى `compactedShuffledFragments` وتسند إليه نتيجة استدعاء `compactFragments` مع مصفوفة `shuffledFragments`.
4. يجب أن تنشئ دالة تسمى `sortFragments` تأخذ مصفوفة من الأجزاء بدون عناصر `undefined` وتعيد مصفوفة جديدة مرتبة حسب الخاصية `id` تصاعديًا، مع الحفاظ على الأجزاء التي تشترك في نفس `id` بترتيبها الأصلي. يجب ألا تستخدم طريقة `sort` المدمجة في JavaScript.
5. يجب أن تعلن متغيرًا يسمى `sortedFragments` وتسند إليه نتيجة استدعاء `sortFragments` مع مصفوفة `compactedShuffledFragments`.
6. يجب أن تنشئ دالة تسمى `dedupeFragments` تأخذ مصفوفة مرتبة من الأجزاء وتعيد مصفوفة جديدة مع إزالة التكرارات، مع الإبقاء على أول ظهور فقط. يجب أن تعرّف التكرارات على أنها جزءان أو أكثر يشتركان في نفس `id`. لكل `id` تتم إزالة تكراره، يجب أن تسجل الدالة رسالة في وحدة التحكم. يجب أن تبدأ الرسالة بالبادئة `[DEDUPED]`.
7. يجب أن تعلن متغيرًا يسمى `dedupedFragments` وتسند إليه نتيجة استدعاء `dedupeFragments` مع مصفوفة `sortedFragments`.
8. يجب أن تنشئ دالة تسمى `fillMissingFragments` تأخذ مصفوفة مرتبة من الأجزاء وتعيد مصفوفة جديدة مع ملء الأجزاء المفقودة بكائنات نائبة. يجب أن تعرّف الأجزاء المفقودة على أنها فجوات في المتتالية بين أصغر وأكبر `id`. يجب أن تكون الكائنات النائبة بالصيغة `{ id: missingId, text: "[...]" }`. لكل كائن نائب يُضاف، يجب أن تسجل الدالة رسالة في وحدة التحكم. يجب أن تبدأ الرسالة بالبادئة `[FILLED]`.
9. يجب أن تعلن متغيرًا يسمى `filledFragments` وتسند إليه نتيجة استدعاء `fillMissingFragments` مع مصفوفة `dedupedFragments`.
10. يجب أن تنشئ دالة تسمى `assembleStory` تأخذ مصفوفة مرتبة من الأجزاء وتعيد سلسلة نصية واحدة تحتوي على جميع نصوص الأجزاء، مفصولة بأسطر جديدة.
11. يجب أن تستخدم `assembleStory` مع `filledFragments` لعرض القصة الكاملة في وحدة التحكم.
12. يجب ألا تغير دوالك `compactFragments`، `sortFragments`، `dedupeFragments`، `fillMissingFragments` و `assembleStory` المصفوفة التي تُستدعى بها.

**مثال**

إليك مثالًا على مصفوفة تحتوي على أجزاء قصة:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

بعد استعادة القصة من `exampleArray`، ستبدو هكذا:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

لا يجب عليك تغيير مصفوفة `shuffledFragments` المملوءة مسبقًا.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

يجب أن يكون لديك دالة `compactFragments`.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` يجب أن تسجل رسالة تبدأ بـ `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

عندما تُستدعى دالة `compactFragments` مع مصفوفة تحتوي على عناصر `undefined`، يجب أن تسجل رسالة في وحدة التحكم تبدأ بـ `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` يجب أن تسجل رسالة تبدأ بـ `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

عندما تُستدعى دالة `compactFragments` مع مصفوفة تحتوي على فراغات، يجب أن تسجل رسالة في وحدة التحكم تبدأ بـ `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

عندما تُستدعى دالة `compactFragments` مع المصفوفة `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`، يجب ألا تسجل أي رسائل في وحدة التحكم تبدأ بـ `[COMPACTED]`.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` يجب أن تُرجع `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

يجب ألا تغير دالة `compactFragments` المصفوفة التي تُمرر إليها.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

يجب أن تكون `compactedShuffledFragments` مصفوفة لا تحتوي على عناصر `undefined`.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

يجب أن يكون لديك دالة `sortFragments`.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` يجب أن تُرجع `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

يجب ألا تستخدم دالة `sortFragments` طريقة `.sort()` المدمجة.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

يجب ألا تغير دالة `sortFragments` المصفوفة التي تُمرر إليها.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

يجب أن يكون لكل عنصر في `sortedFragments` قيمة `id` أقل من أو تساوي قيمة `id` للعنصر التالي.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

يجب أن يكون لديك دالة `dedupeFragments`.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` يجب أن تُرجع `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` يجب أن تسجل رسالة `[DEDUPED]` لكل `id` مكرر.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` يجب أن تسجل رسالة `[DEDUPED]` عندما يكون آخر `id` مكررًا.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

عندما تُستدعى دالة `dedupeFragments` مع مصفوفة تحتوي على تكرارات، يجب أن تسجل رسالة تبدأ بـ `[DEDUPED]` لكل `id` مكرر.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

يجب ألا تغير دالة `dedupeFragments` المصفوفة التي تُمرر إليها.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

يجب ألا تحتوي `dedupedFragments` على قيم `id` مكررة.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

يجب أن يكون لديك دالة `fillMissingFragments`.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` يجب أن تُرجع `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` يجب أن تسجل رسالة تبدأ بـ `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

عندما تضيف دالة `fillMissingFragments` عنصرًا نائبًا، يجب أن تسجل رسالة تبدأ بـ `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

يجب ألا تغير دالة `fillMissingFragments` المصفوفة التي تُمرر إليها.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

يجب ألا تحتوي `filledFragments` على فجوات في متتالية `id`.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

يجب أن يكون لديك دالة `assembleStory`.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` يجب أن تُرجع السلسلة النصية `"Hello\nWorld"`.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

يجب ألا تغير دالة `assembleStory` المصفوفة التي تُمرر إليها.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

يجب أن تستدعي `console.log` مع نتيجة `assembleStory(filledFragments)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
