---
id: 671fa22e99042b1e91fe9d2e
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Inside the `.btn-container`, create three buttons for each type of story. Dê a cada botão uma classe `btn` e um `id` de `scary-btn`, `funny-btn` e `adventure-btn`. O texto para cada botão deve ser `Scary Story`, `Funny Story` e `Adventure Story`, respectivamente.

# --hints--

Dentro do elemento `.btn-container`, você deve ter três botões.

```js
assert.lengthOf(document.querySelectorAll('.btn-container > button') ,3);
```

Todos os três botões devem ter uma classe `btn`.

```js
assert.equal(document.querySelectorAll('.btn-container > .btn').length, 3);
```

Você deve ter um botão com o `id` de `scary-btn` e um texto de `Scary Story`.

```js
const scaryBtn = document.querySelector('#scary-btn');
assert.exists(scaryBtn);
assert.equal(scaryBtn.textContent.trim(), 'Scary Story');
```

Você deve ter um botão com o `id` `funny-btn` e um texto `Funny Story`.

```js
const funnyBtn = document.querySelector('#funny-btn');
assert.exists(funnyBtn);
assert.equal(funnyBtn.textContent.trim(), 'Funny Story');
```

Você deve ter um botão com o `id` `adventure-btn` e um texto `Adventure Story`.

```js
const adventureBtn = document.querySelector('#adventure-btn');
assert.exists(adventureBtn);
assert.equal(adventureBtn.textContent.trim(), 'Adventure Story');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Storytelling App</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Want to hear a short story?</h1>

    <main class="story-container">
--fcc-editable-region--
        <div class="btn-container">
                
        </div>
--fcc-editable-region--
</main>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --dark-grey: #1b1b32;
    --black: #000;
    --white: #fff;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
}

body {
    background-color: var(--dark-grey);
    color: var(--white);
}

h1,
#result {
    text-align: center;
}

h1 {
    margin: 10px 0 15px;
}

.story-container {
    margin: auto;
    padding: 10px;
    width: 80%;
    border-style: double;
    border-width: 14px;
    border-color: var(--white);
}

.btn-container {
    display: flex;
    flex-direction: column;
    justify-content: center;
    align-items: center;
}

@media (min-width: 760px) {
    .btn-container {
        flex-direction: row;
    }
}

#result {
    margin-top: 15px;
    font-size: 1.2rem;
    line-height: 30px;
}

.btn {
    cursor: pointer;
    width: 200px;
    margin: 10px 0 10px 0.5rem;
    color: var(--black);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}
```
