---
id: 615f34ecc1091b4fd5a8a484
title: 步驟4
challengeType: 0
dashedName: step-4
---

# --description--

Within your `head` element, add a `link` element with the `rel` attribute set to `stylesheet` and the `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

這將導入 `Open Sans` 字體系列，字體粗細值爲 `400`、`700` 和 `800`。

還要添加一個 `link` 元素來鏈接 `styles.css` 文件。

# --hints--

你的代碼中應該有兩個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Your `link` elements should be void elements, they should not have an end tag `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

你的兩個 `link` 元素應該在 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Both of your `link` elements should have the `rel` attribute set to `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

你的其中一個 `link` 元素的 `href` 屬性應設置爲 `styles.css`。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

One of your `link` elements should have an `href` attribute set to `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>

</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
--fcc-editable-region--
```

```css

```
