---
id: 655cd899f8de09431eabb40c
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Merge sort é um algoritmo de ordenação que usa o princípio de dividir e conquistar para ordenar coleções de dados. Ou seja, ele "divide" uma coleção em subpartes menores e "conquista" as subpartes ordenando-as independentemente, depois mescla as subpartes ordenadas.

Comece criando uma função chamada `merge_sort` que recebe o parâmetro `array`.Esta função irá lidar com a tarefa de ordenar uma lista de números.Use a palavra-chave `pass` no corpo da função.

# --hints--

Você deve declarar uma função chamada `merge_sort` com um único parâmetro: `array`. Não esqueça de usar a palavra-chave `pass` no corpo da função.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
