---
id: 6564aee9c077774ea49c3faf
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

O algoritmo de merge sort realiza três ações:

- Divida uma sequência não ordenada de itens em subpartes
- Ordene os itens nas subpartes
- Mesclar as subpartes ordenadas

O acima acontece recursivamente até que as subpartes sejam mescladas na sequência ordenada completa.Vamos começar dividindo a sequência.

Primeiro, substitua a palavra-chave `pass` por uma variável `middle_point`.Use o operador de divisão inteira (`//`) para dividir o comprimento da lista `array` por 2 e atribuir o resultado à sua nova variável `middle_point`.Lembre-se de indentar seu código.

# --hints--

Você não deve ter `pass` no corpo da função.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

Você deve declarar uma variável `middle_point` e atribuir `len(array) // 2` a ela.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
