---
id: 656627b47bd2d2a4afbc945d
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

No passo anterior, você obteve o ponto médio.Você pode usá-lo para dividir `array` em duas partes e atribuir cada parte a novas variáveis. 

Use a sintaxe de slice para extrair a metade esquerda de `array` e atribuí-la a uma variável chamada `left_part`.

# --hints--

Você deve ter uma variável chamada `left_part` na sua função `merge_sort`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Você deve usar a sintaxe de slice para extrair a metade esquerda do array e atribuí-la à variável `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
