---
id: 656639771fed09b5c0e8fe71
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora que você dividiu a lista `array` em duas listas separadas, você continuará dividindo cada lista até que cada elemento esteja sozinho em sua própria lista.Uma lista com um único número está sempre ordenada.

Para fazer isso, chame recursivamente `merge_sort` dentro da sua função e passe `left_part` como argumento para a chamada.

# --hints--

Você deve chamar `merge_sort` com o argumento `left_part` na parte inferior do corpo da sua função.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
