---
id: 656680b0fc79f2c38a34d90e
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Agora é hora de ordenar e mesclar as listas (`left_part` e `right_part`) na lista original. 

Você pode fazer isso comparando os elementos em ambas as listas e mesclando o elemento menor à lista principal.Você vai fazer essa comparação para todos os índices em `left_part` e `right_part`.

Crie três variáveis: `left_array_index`, `right_array_index` e `sorted_index` e defina seus valores como `0`.Essas variáveis vão ajudar você a acompanhar cada índice durante o processo de ordenação.

# --hints--

Você deve ter uma variável chamada `left_array_index` dentro da sua função `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

O valor de `left_array_index` deve ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

Você deve ter uma variável chamada `right_array_index` dentro da sua função `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

O valor de `right_array_index` deve ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

Você deve ter uma variável chamada `sorted_index` dentro da sua função `merge_sort`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

O valor de `sorted_index` deve ser `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    merge_sort(left_part)
    merge_sort(right_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
