---
id: 656702f8b4cbd8cbf0a433c6
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Em seguida, você precisa criar um loop que continue enquanto houver elementos restantes em `left_part` e `right_part`.

Para isso, crie um loop `while` com duas condições: uma que verifica se o `left_array_index` é menor que o comprimento de `left_part` e outra condição que verifica se o `right_array_index` é menor que o comprimento de `right_part`.

# --hints--

Seu loop `while` deve verificar se `left_array_index` é menor que o comprimento de `left_part` e se `right_array_index` é menor que o comprimento de `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Use a palavra-chave `pass` no corpo do loop.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
