---
id: 656707e11e671ed54c96f7ec
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Quando a condição `if` avalia para `True`, isso significa que o elemento na lista `left_part` é menor do que o elemento com o qual está sendo comparado na lista `right_part`.

Nesse caso, você pode colocar o elemento encontrado em `left_array_index` em `left_part` dentro do array ordenado.

Dentro do bloco `if`, remova `pass` e atribua `left_part[left_array_index]` a `array[sorted_index]`.

# --hints--

Você deve substituir `pass` por `array[sorted_index] = left_part[left_array_index]` no corpo da instrução `if`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
