---
id: 656710d1e0ec62253426db24
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

O loop `while` que você criou compara um elemento de `left_part` com outro em `right_part`, depois adiciona o elemento menor à lista principal `array`.

Ele continuará essa operação até que não haja mais elementos para serem comparados. Mas `left_part` ainda pode ter elementos restantes enquanto `right_part` não tem nenhum e vice-versa.

Crie outro loop `while` que execute quando `left_array_index` for menor que `len(left_part)`. Nos próximos passos, você usará isso para copiar os elementos restantes em `left_part` para a lista `array`.

# --hints--

Seu loop `while` deve ter esta condição: `left_array_index < len(left_part)`. Não esqueça de usar a palavra-chave `pass` no corpo do loop.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

Você deve ter a palavra-chave `pass` no corpo do seu loop `while`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--
    
--fcc-editable-region--
```
