---
id: 6569b4e0bd29d17d4b53b16c
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Agora, você vai replicar a mesma lógica do loop `while` para `right_part`.

Crie um loop `while` que execute quando `right_array_index` for menor que `len(right_part)`. Dentro do loop `while`, atribua `right_part[right_array_index]` a `array[sorted_index]`. No final, incremente `right_array_index` e `sorted_index` em `1`.

# --hints--

Seu loop `while` deve verificar se `right_array_index` é menor que `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

Você deve atribuir `right_part[right_array_index]` a `array[sorted_index]` no corpo do loop. 

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

Você deve incrementar `right_array_index` em `1` dentro do loop `while`. Note que o incremento deve ser feito após atribuir o elemento da parte direita ao array ordenado.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

Você deve incrementar o valor atual de `sorted_index` em `1` dentro do loop `while`. Note que o incremento deve ser feito após atribuir o elemento da parte direita ao array ordenado.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
