---
id: 6577254891048c8f2c19e961
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Na parte inferior do corpo da sua função, chame a função `merge_sort()` novamente.Desta vez, passe `right_part` como argumento para a chamada da função.

# --hints--

Você deve chamar `merge_sort()` com o argumento `right_part` no final do corpo da sua função.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
