---
id: 657f59751d5a5c9b069ae0f3
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Como você aprendeu em uma lição anterior, você pode usar a variável `__name__` para determinar se um script Python está sendo executado como o programa principal ou se está sendo importado como um módulo (código escrito em outro arquivo Python).

Se o valor de `__name__` estiver definido como `'__main__'`, isso implica que o script atual é o programa principal e não um módulo.

Neste projeto, você usará o script atual como o programa principal.

Crie uma declaração `if` que verifique se o valor de `__name__` é `'__main__'`. Dentro da declaração `if`, crie uma lista chamada `numbers` e atribua a ela estes valores: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

Sua declaração `if` deve verificar se `__name__` tem um valor de `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

Você deve ter uma lista `numbers` com esses valores: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
