---
id: 657f3a661730891aa64f3568
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Use a sintaxe de slice para extrair a metade direita de `array` e atribuí-la a uma variável chamada `right_part`.

# --hints--

Você deve ter uma variável chamada `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Você deve usar a sintaxe de slice para extrair a metade direita do array e atribuí-la à variável `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
