---
id: 5e44414f903586ffb414c950
title: بناء مشروع آلة حاسبة الاحتمالات
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462364
dashedName: build-a-probability-calculator-project
---

# --description--

افترض وجود قبعة تحتوي على 5 كرات زرقاء، و4 كرات حمراء، و2 كرة خضراء. ما احتمال أن تحتوي سحبة عشوائية من 4 كرات على كرة حمراء واحدة على الأقل وكراتين خضراء؟ رغم إمكانية حساب الاحتمال باستخدام الرياضيات المتقدمة، فإن الطريقة الأسهل هي كتابة برنامج لإجراء عدد كبير من التجارب لتقدير احتمال تقريبي.

في هذا المشروع، ستكتب برنامجًا لتحديد الاحتمال التقريبي لسحب كرات معينة عشوائيًا من قبعة.

أولًا، أنشئ صنف `Hat` في `main.py`. يجب أن يأخذ الصنف عددًا متغيرًا من المعلمات التي تحدد عدد الكرات من كل لون في القبعة. على سبيل المثال، يمكن إنشاء كائن من الصنف بأي من الطرق التالية:

```py
hat1 = Hat(yellow=3, blue=2, green=6)
hat2 = Hat(red=5, orange=4)
hat3 = Hat(red=5, orange=4, black=1, blue=0, pink=2, striped=9)
```

سيتم إنشاء القبعة دائمًا مع كرة واحدة على الأقل. يجب تحويل المعلمات الممررة إلى كائن القبعة عند الإنشاء إلى متغير مثيل `contents`. يجب أن يكون `contents` قائمة من السلاسل النصية تحتوي على عنصر واحد لكل كرة في القبعة. كل عنصر في القائمة يجب أن يكون اسم لون يمثل كرة واحدة من ذلك اللون. على سبيل المثال، إذا كانت القبعة `{'red': 2, 'blue': 1}`، يجب أن تكون `contents` هي `['red', 'red', 'blue']`.

يجب أن يحتوي صنف `Hat` على طريقة `draw` تقبل معلمة تحدد عدد الكرات التي يجب سحبها من القبعة. يجب أن تزيل هذه الطريقة كرات عشوائية من `contents` وتُرجع تلك الكرات كقائمة من السلاسل النصية. لا يجب إعادة الكرات إلى القبعة أثناء السحب، مثل تجربة جرة بدون استبدال. إذا كان عدد الكرات المطلوب سحبها أكبر من الكمية المتاحة، أرجع كل الكرات.

بعد ذلك، أنشئ دالة `experiment` في `main.py` (ليست داخل صنف `Hat`). يجب أن تقبل هذه الدالة المعلمات التالية:

- `hat`: كائن قبعة يحتوي على كرات يجب نسخه داخل الدالة.
- `expected_balls`: كائن يحدد المجموعة الدقيقة من الكرات التي يجب محاولة سحبها من القبعة في التجربة. على سبيل المثال، لتحديد احتمال سحب كرتين زرقاوين وكرة حمراء واحدة من القبعة، عيّن `expected_balls` إلى `{'blue':2, 'red':1}`.
- `num_balls_drawn`: عدد الكرات التي يجب سحبها من القبعة في كل تجربة.
- `num_experiments`: عدد التجارب التي يجب إجراؤها. (كلما زاد عدد التجارب، كان الاحتمال التقريبي أدق.)

يجب أن تُرجع دالة `experiment` احتمالًا.

على سبيل المثال، إذا أردت تحديد احتمال الحصول على كرتين حمراوين على الأقل وكرة خضراء واحدة عند سحب خمس كرات من قبعة تحتوي على ست كرات سوداء وأربع حمراء وثلاث خضراء. للقيام بذلك، ستُجري `N` تجربة، تحصي عدد المرات `M` التي تحصل فيها على كرتين حمراوين على الأقل وكرة خضراء واحدة، وتقدر الاحتمال كـ `M/N`. تتكون كل تجربة من البدء بقبعة تحتوي على الكرات المحددة، سحب عدة كرات، والتحقق مما إذا حصلت على الكرات التي حاولت سحبها.

إليك كيفية استدعاء دالة `experiment` بناءً على المثال أعلاه مع 2000 تجربة:

```py
hat = Hat(black=6, red=4, green=3)
probability = experiment(hat=hat,
                  expected_balls={'red':2,'green':1},
                  num_balls_drawn=5,
                  num_experiments=2000)
```

سيكون الناتج شيئًا مثل هذا:

```bash
0.356
```

بما أن هذا يعتمد على السحبات العشوائية، سيكون الاحتمال مختلفًا قليلاً في كل مرة تُشغل فيها الشفرة.

_تلميح: فكر في استخدام الوحدات التي تم استيرادها مسبقًا في الأعلى. لا تهيئ البذرة العشوائية داخل الملف._

ملاحظة: افتح وحدة تحكم المتصفح بالضغط على F12 لرؤية إخراج أكثر تفصيلاً للاختبارات.

# --hints--

يجب أن تضيف عملية إنشاء كائن `hat` المحتويات الصحيحة.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_class_contents(self):
        hat = probability_calculator.Hat(red=3,blue=2)
        actual = hat.contents
        expected = ["red","red","red","blue","blue"]
        self.assertEqual(actual, expected, 'Expected creation of hat object to add correct contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

يجب أن تُنقص طريقة `draw` في صنف `hat` عدد العناصر في القائمة contents.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)

class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw(self):
        hat = probability_calculator.Hat(red=5,blue=2)
        actual = hat.draw(2)
        expected = ['blue', 'red']
        self.assertEqual(actual, expected, 'Expected hat draw to return two random items from hat contents.')
        actual = len(hat.contents)
        expected = 5
        self.assertEqual(actual, expected, 'Expected hat draw to reduce number of items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

يجب أن تتصرف طريقة `draw` بشكل صحيح عندما يكون عدد الكرات المطلوب سحبها أكبر من عدد الكرات في القبعة.

```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_hat_draw_2(self):
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        actual = sorted(hat.draw(20))
        expected = sorted(['yellow', 'yellow', 'yellow', 'yellow', 'yellow', 'red', 'green', 'green', 'green', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'blue', 'test'])
        self.assertEqual(actual, expected, 'Expected hat draw to return all items from hat contents.')
        actual = len(hat.contents)
        expected = 0
        self.assertEqual(actual, expected, 'Expected hat draw to leave no items in contents.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

يجب أن تُرجع طريقة `experiment` احتمالًا مختلفًا.


```js
({
  test: () => {
    pyodide.FS.writeFile("/home/pyodide/probability_calculator.py", code);
    pyodide.FS.writeFile(
      "/home/pyodide/test_module.py",
      `
import unittest
import probability_calculator
from importlib import reload

reload(probability_calculator)

probability_calculator.random.seed(95)
class UnitTests(unittest.TestCase):
    maxDiff = None
    def test_prob_experiment(self):
        hat = probability_calculator.Hat(blue=3,red=2,green=6)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"blue":2,"green":1}, num_balls_drawn=4, num_experiments=1000)
        actual = probability
        expected = 0.272
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        hat = probability_calculator.Hat(yellow=5,red=1,green=3,blue=9,test=1)
        probability = probability_calculator.experiment(hat=hat, expected_balls={"yellow":2,"blue":3,"test":1}, num_balls_drawn=20, num_experiments=100)
        actual = probability
        expected = 1.0
        self.assertAlmostEqual(actual, expected, delta = 0.01, msg = 'Expected experiment method to return a different probability.')
        `
    );
    const testCode = `
from unittest import main
import test_module
from importlib import reload

reload(test_module)
t = main(module='test_module', exit=False)
t.result.wasSuccessful()
`;
    const out = runPython(testCode);
    assert(out);
  },
});
```

# --seed--

## --seed-contents--

```py
import copy
import random

class Hat:
    pass

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):
    pass
```

# --solutions--

```py
import copy
import random

class Hat:
    def __init__(self, **hat):
        self.hat = hat
        contents = []
        for i in hat:
            for j in range(hat[i]):
                contents.append(i)           
        self.contents = contents
               

    def draw(self, number): 
        drawn = []
        if number >= len(self.contents):
            drawn.extend(self.contents)
            self.contents = []
        else:
            for i in range(number):
                drawn.append(
                    self.contents.pop(random.randrange(len(self.contents)))
                )                
        return drawn

def experiment(hat, expected_balls, num_balls_drawn, num_experiments):    
    expected_balls_list = []
    drawn_list = []
    success = 0
    for i in expected_balls:
        for j in range(expected_balls[i]):
            expected_balls_list.append(i)
    for j in range(num_experiments):
        hat_copy = copy.deepcopy(hat)
        drawn_list.append(hat_copy.draw(num_balls_drawn))        
        exp_ball_list_copy = expected_balls_list[:]
        for k in range(len(drawn_list[j])):
            try:
                ind = exp_ball_list_copy.index(drawn_list[j][k])
                exp_ball_list_copy.pop(ind)
            except:
                continue
        if len(exp_ball_list_copy) == 0:
            success += 1        
    probability = success/num_experiments    
    return probability
```
