---
id: 5ddb965c65d27e1512d44d9a
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個專案中，你將學習如何創建一個熱量計數器表單，該表單可讓使用者輸入他們每日的熱量預算以及各種餐點的熱量數值。該表單將計算並顯示使用者是處於熱量赤字還是過剩。

你已經獲得樣板的 CSS 和 HTML。然而，你需要建置你的熱量計算器表單。

隨意探索 HTML 和 CSS，然後添加一個 `form` 元素並給它一個設定為 `calorie-counter` 的 `id`。

# --hints--

你應該有一個 `form` 元素。

```js
assert.exists(document.querySelector('form'));
```

你的 `form` 元素應該有一個設定為 `calorie-counter` 的 `id` 屬性。

```js
assert.equal(document.querySelector('form')?.id, 'calorie-counter');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
--fcc-editable-region--

--fcc-editable-region--
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
