---
id: 63b607af6fcdb119aae9b16a
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

在你的表單中，使用者應該能夠為各種餐點型別添加熱量數值。

創建一個 `fieldset` 元素，並將 `id` 設定為 `breakfast`。

在該元素內，創建一個帶有文字 `Breakfast` 的 `legend`，以及一個 `class` 設為 `input-container` 的空 `div`。

# --hints--

你應該在你的 `form` 中創建一個 `fieldset` 元素。

```js
assert.exists(document.querySelector('form fieldset'));
```

你的 `fieldset` 元素應該放在你的 `input` 元素之後。

```js
assert.equal(document.querySelector('form fieldset')?.previousElementSibling?.tagName, "INPUT");
```

你的 `fieldset` 元素應該設定一個 `id` 為 `breakfast`。

```js
assert.equal(document.querySelector('form fieldset')?.id, "breakfast");
```

你應該在你的 `fieldset` 中創建一個 `legend` 元素。

```js
assert.exists(document.querySelector('form fieldset legend'));
```

你的 `legend` 元素應該有文字 `Breakfast`。

```js
assert.equal(document.querySelector('form fieldset legend')?.innerText, "Breakfast");
```

你應該在你的 `fieldset` 中創建一個 `div` 元素。

```js
assert.exists(document.querySelector('form fieldset div'));
```

你的 `div` 元素應該位於你的 `legend` 元素之後。

```js
assert.equal(document.querySelector('form fieldset div')?.previousElementSibling?.tagName, "LEGEND");
```

你的 `div` 元素應該設定 `class` 為 `input-container`。

```js
assert.equal(document.querySelector('form fieldset div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
