---
id: 587d824a367417b2b2512c46
title: Impara come funzionano le asserzioni in JavaScript
challengeType: 2
forumTopicId: 301589
dashedName: learn-how-javascript-assertions-work
---

# --description--

Per affrontare queste sfide dovrai scrivere il tuo codice usando uno dei seguenti metodi:

- Clona <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">questo repository GitHub</a> e completa queste sfide in locale.
- Usa un generatore di siti a tua scelta per completare il progetto. Assicurati di includere tutti i file dal nostro repository GitHub.

# --instructions--

All’interno di `tests/1_unit-tests.js` sotto il test etichettato `#1` nella suite `Basic Assertions`, cambia ogni `assert` in `assert.isNull` o `assert.isNotNull` per far passare il test (dovrebbe valutarsi a `true`). Non modificare gli argomenti passati alle asserzioni.

# --hints--

Tutti i test dovrebbero passare.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `isNull` vs. `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[0].method, 'isNull', 'Null is null');
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `isNull` vs. `isNotNull`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=0');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[1].method, 'isNotNull', '1 is not null');
```
