---
id: 587d824d367417b2b2512c51
title: Verifica se un array contiene un elemento
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#12` nella suite `Arrays`, cambia ogni `assert` in `assert.include` o `assert.notInclude` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati agli assert.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `include` vs. `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a back-end language !!'
);
```
