---
id: 587d824b367417b2b2512c4a
title: "usa il doppio uguale per verificare l'uguaglianza"
challengeType: 2
forumTopicId: 301609
dashedName: use-the-double-equals-to-assert-equality
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`equal()` confronta gli oggetti usando `==`.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#5` nella suite `Equality`, cambia ogni `assert` in `assert.equal` o `assert.notEqual` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati agli assert.

# --hints--

Tutti i test dovrebbero passare.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `equal` vs. `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'equal',
    'Numbers are coerced into strings with == '
  );
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `equal` vs. `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notEqual',
    ' == compares object references'
  );
```

Dovresti scegliere il metodo corretto per la terza asserzione - `equal` vs. `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'equal',
    "6 * '2' is 12 ! It should be equal to '12'"
  );
```

Dovresti scegliere il metodo corretto per la quarta asserzione - `equal` vs. `notEqual`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=4');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.assertions[3].method, 'notEqual', "6 + '2' is '62'...");
```
