---
id: 587d824f367417b2b2512c5b
title: Esegui test funzionali su una risposta API usando Chai-HTTP IV - metodo PUT
challengeType: 2
forumTopicId: 301591
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iv---put-method
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Questo esercizio è simile al precedente.

Ora che sai come testare una richiesta `PUT`, tocca a te farlo da zero.

# --instructions--

All'interno di `tests/2_functional-tests.js`, modifica il test `'Send {surname: "da Verrazzano"}'` (`// #4`) e usa i metodi `put` e `send` per testare l'endpoint `'/travellers'`.

Invia il seguente oggetto JSON con la tua richiesta PUT:

```json
{
  "surname": "da Verrazzano"
}
```

Controlla quanto segue all'interno della callback `request.end`:

1.  Il `status` dovrebbe essere `200`
2.  Il `type` dovrebbe essere `application/json`
3.  Il `body.name` dovrebbe essere `Giovanni`
4.  Il `body.surname` dovrebbe essere `da Verrazzano`

Segui l'ordine delle asserzioni sopra indicato - ci affidiamo a questo. Inoltre, assicurati di rimuovere `assert.fail()` una volta completato.

# --hints--

Tutti i test dovrebbero superare

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti testare che `res.status` sia 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Dovresti testare che `res.type` sia `'application/json'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Dovresti testare che `res.body.name` sia `'Giovanni'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Giovanni\1/);
```

Dovresti testare che `res.body.surname` sia `'da Verrazzano'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=3');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")da Verrazzano\1/);
```
