---
id: 587d824c367417b2b2512c4e
title: Verifica se un valore è inferiore o almeno pari a un altro
challengeType: 2
forumTopicId: 301606
dashedName: test-if-one-value-is-below-or-at-least-as-large-as-another
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#9` nella suite `Comparisons`, cambia ciascun `assert` in `assert.isBelow` o `assert.isAtLeast` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati alle asserzioni.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `isBelow` vs. `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isAtLeast',
  '5 is at least (>=) 5'
);
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `isBelow` vs. `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isAtLeast',
  '2 * Math.random() is at least 0'
);
```

Dovresti scegliere il metodo corretto per la terza asserzione - `isBelow` vs. `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isBelow', '1 is smaller than 2');
```

Dovresti scegliere il metodo corretto per la quarta asserzione - `isBelow` vs. `isAtLeast`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=8');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[3].method,
  'isBelow',
  '2/3 (0.6666) is smaller than 1'
);
```
