---
id: 587d824d367417b2b2512c52
title: Verifica se un valore è una stringa
challengeType: 2
forumTopicId: 301599
dashedName: test-if-a-value-is-a-string
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`isString` o `isNotString` verifica che il valore effettivo sia una stringa.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#13` nella suite `Strings`, modifica ciascun `assert` in `assert.isString` o `assert.isNotString` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati agli assert.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `isString` vs. `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'isNotString',
  'A float number is not a string'
);
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `isString` vs. `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'isString',
  'environment vars are strings (or undefined)'
);
```

Dovresti scegliere il metodo corretto per la terza asserzione - `isString` vs. `isNotString`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=12');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'isString', 'A JSON is a string');
```
