---
id: 587d824d367417b2b2512c51
title: Verificare se un array contiene un elemento
challengeType: 2
forumTopicId: 301603
dashedName: test-if-an-array-contains-an-item
---

# --description--

As a reminder, this project is being built upon the following starter project cloned from <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

All'interno di `tests/1_unit-test.js` sotto il test etichettato `#12` nella suite `Arrays` , cambiare ogni `assert` in `assert.include` o `assert.notInclude` per far passare il test (dovrebbe risultare `true`). Non alterare gli argomenti passati alle asserzioni.

# --hints--

Tutti i test dovrebbero essere superati.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `include` oppure `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'notInclude',
  "It's summer in july..."
);
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `include` oppure `notInclude`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=11');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'include',
  'JavaScript is a backend language !!'
);
```

