---
id: 587d824f367417b2b2512c5a
title: Esegui test funzionali su una risposta API usando Chai-HTTP III - metodo PUT
challengeType: 2
forumTopicId: 301590
dashedName: run-functional-tests-on-an-api-response-using-chai-http-iii---put-method
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Quando testi una `PUT` richiesta, spesso invii dati insieme ad essa. I dati che includi nella tua `PUT` richiesta si chiamano corpo della richiesta.

Per inviare una `PUT` richiesta e un oggetto JSON all'endpoint `'/travellers'`, puoi usare i metodi `chai-http` e `put` del plugin `send`:

```js
chai
  .request(server)
  .keepOpen()
  .put('/travellers')
  .send({
    "surname": [last name of a traveller of the past]
  })
  ...
```

E la rotta risponde con:

```json
{
  "name": [first name],
  "surname": [last name],
  "dates": [birth - death years]
}
```

Consulta il codice del server per le diverse risposte all'endpoint `'/travellers'`.

# --instructions--

All'interno di `tests/2_functional-tests.js`, modifica il test `'Send {surname: "Colombo"}'` (`// #3`) e usa i metodi `put` e `send` per testare l'endpoint `'/travellers'`.

Invia il seguente oggetto JSON con la tua richiesta PUT:

```json
{
  "surname": "Colombo"
}
```

Controlla quanto segue all'interno della callback `request.end`:

1.  Il `status` dovrebbe essere `200`
2.  Il `type` dovrebbe essere `application/json`
3.  Il `body.name` dovrebbe essere `Cristoforo`
4.  Il `body.surname` dovrebbe essere `Colombo`

Segui l'ordine delle asserzioni sopra indicato — ci affidiamo a questo. Inoltre, assicurati di rimuovere `assert.fail()` una volta completato.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti testare che `res.status` sia 200.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Dovresti testare che `res.type` sia `'application/json'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.type');
assert.match(data.assertions[1].args[1], /('|")application\/json\1/);
```

Dovresti testare che `res.body.name` sia `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'equal');
assert.equal(data.assertions[2].args[0], 'res.body.name');
assert.match(data.assertions[2].args[1], /('|")Cristoforo\1/);
```

Dovresti testare che `res.body.surname` sia `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=2');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'equal');
assert.equal(data.assertions[3].args[0], 'res.body.surname');
assert.match(data.assertions[3].args[1], /('|")Colombo\1/);
```
