---
id: 587d8250367417b2b2512c5d
title: Esegui test funzionali usando un browser headless
challengeType: 2
forumTopicId: 301595
dashedName: run-functional-tests-using-a-headless-browser
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Nella pagina c'è un modulo di input. Invia i dati all'endpoint `PUT /travellers` come richiesta AJAX.

Quando la richiesta si completa con successo, il codice client aggiunge un `<div>` contenente le informazioni nella risposta al DOM.

Ecco un esempio di come usare Zombie.js per interagire con il modulo:

```js
test('Submit the surname "Polo" in the HTML form', function (done) {
  browser.fill('surname', 'Polo').then(() => {
    browser.pressButton('submit', () => {
      browser.assert.success();
      browser.assert.text('span#name', 'Marco');
      browser.assert.text('span#surname', 'Polo');
      browser.assert.elements('span#dates', 1);
      done();
    });
  });
});
```

Per prima cosa, il metodo `fill` dell'oggetto `browser` riempie il campo `surname` del modulo con il valore `'Polo'`. `fill` restituisce una promise, quindi `then` viene concatenato a essa.

All'interno della callback `then`, il metodo `pressButton` dell'oggetto `browser` viene usato per invocare il listener dell'evento `submit` del modulo. Il metodo `pressButton` è asincrono.

Poi, una volta ricevuta la risposta dalla richiesta AJAX, vengono fatte alcune asserzioni per confermare:

1.  Lo stato della risposta è `200`
2.  Il testo all'interno dell'elemento `<span id='name'></span>` corrisponde a `'Marco'`
3.  Il testo all'interno dell'elemento `<span id='surname'></span>` corrisponde a `'Polo'`
4.  Esiste `1` elemento `<span id='dates'></span>`

Infine, viene invocata la callback `done`, necessaria a causa del test asincrono.

# --instructions--

All'interno di `tests/2_functional-tests.js`, nel test `'Submit the surname "Colombo" in the HTML form'` (`// #5`), automatizza quanto segue:

1.  Compila il modulo con il cognome `Colombo`
2.  Premi il pulsante di invio

E all'interno della callback `pressButton`:

1.  Asserisci che lo stato sia OK `200`
2.  Asserisci che il testo dentro l'elemento `span#name` sia `'Cristoforo'`
3.  Asserisci che il testo dentro l'elemento `span#surname` sia `'Colombo'`
4.  Asserisci che l'elemento/ gli elementi `span#dates` esistano e che la loro quantità sia `1`

Non dimenticare di rimuovere la chiamata `assert.fail()`.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti asserire che la richiesta del browser headless sia andata a buon fine.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Dovresti asserire che il testo dentro l'elemento `span#name` sia `'Cristoforo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Cristoforo\1/);
```

Dovresti asserire che il testo dentro l'elemento `span#surname` sia `'Colombo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Colombo\1/);
```

Dovresti asserire che l'elemento `span#dates` esista e che la sua quantità sia 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=5');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
