---
id: 587d824b367417b2b2512c47
title: Verifica se una variabile o funzione è definita
challengeType: 2
forumTopicId: 301602
dashedName: test-if-a-variable-or-function-is-defined
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#2` nella suite `Basic Assertions`, cambia ciascun `assert` in `assert.isDefined()` o `assert.isUndefined()` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati alle asserzioni.

# --hints--

Tutti i test dovrebbero passare.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `isDefined` vs. `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'isDefined',
    'Null is not undefined'
  );
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `isDefined` vs. `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'isUndefined',
    'Undefined is undefined'
  );
```

Dovresti scegliere il metodo corretto per la terza asserzione - `isDefined` vs. `isUndefined`.

```js
  const response = await fetch(code + '/_api/get-tests?type=unit&n=1');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[2].method,
    'isDefined',
    'A string is not undefined'
  );
```
