---
id: 587d824e367417b2b2512c58
title: Esegui test funzionali sugli endpoint API usando chai-http
challengeType: 2
forumTopicId: 301593
dashedName: run-functional-tests-on-api-endpoints-using-chai-http
---

# --description--

Come promemoria, questo progetto si basa sul progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Mocha ti permette di testare operazioni asincrone come le chiamate agli endpoint API con un plugin chiamato `chai-http`.

Di seguito un esempio di test che usa `chai-http` per una suite chiamata `'GET /hello?name=[name] => "hello [name]"'`:

```js
suite('GET /hello?name=[name] => "hello [name]"', function () {
  test('?name=John', function (done) {
    chai
      .request(server)
      .keepOpen()
      .get('/hello?name=John')
      .end(function (err, res) {
        assert.equal(res.status, 200, 'Response status should be 200');
        assert.equal(res.text, 'hello John', 'Response should be "hello John"');
        done();
      });
  });
});
```

Il test invia una richiesta `GET` al server con un nome come stringa di query URL (`?name=John`). Nella funzione di callback del metodo `end`, viene ricevuto l’oggetto risposta (`res`) che contiene la proprietà `status`.

Il primo `assert.equal` verifica se lo stato è uguale a `200`. Il secondo `assert.equal` verifica che la stringa di risposta (`res.text`) sia uguale a `"hello John"`.

Nota anche il parametro `done` nella funzione di callback del test. Chiamarlo senza argomenti alla fine di un test è necessario per segnalare che l’operazione asincrona è completata.

Infine, nota il metodo `keepOpen` subito dopo il metodo `request`. Normalmente eseguirai i test da riga di comando o come parte di un processo di integrazione automatizzato, e potresti lasciare che `chai-http` avvii e fermi automaticamente il tuo server.

Tuttavia, i test che vengono eseguiti quando invii il collegamento al tuo progetto richiedono che il server sia attivo, quindi devi usare il metodo `keepOpen` per impedire che `chai-http` fermi il server.

# --instructions--

All’interno di `tests/2_functional-tests.js`, modifica il test `'Test GET /hello with no name'` (`// #1`) per verificare `status` e `text` della risposta in modo che il test passi. Non modificare gli argomenti passati agli assert.

Non deve esserci alcuna query URL. Senza una query URL con nome, l’endpoint risponde con `hello Guest`.

# --hints--

Tutti i test devono passare

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti testare che `res.status` == 200

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'equal');
assert.equal(data.assertions[0].args[0], 'res.status');
assert.equal(data.assertions[0].args[1], '200');
```

Dovresti testare che `res.text` == `'hello Guest'`

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=0');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'equal');
assert.equal(data.assertions[1].args[0], 'res.text');
assert.match(data.assertions[1].args[1], /('|")hello Guest\1/);
```
