---
id: 587d824d367417b2b2512c54
title: Usare espressioni regolari per testare una stringa
challengeType: 2
forumTopicId: 301608
dashedName: use-regular-expressions-to-test-a-string
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`match()` verifica che il valore effettivo corrisponda all'espressione regolare passata come secondo argomento.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#15` nella suite `Strings`, cambia ciascun `assert` in `assert.match` o `assert.notMatch` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati agli assert.

# --hints--

Tutti i test dovrebbero passare.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti scegliere il metodo corretto per la prima asserzione - `match` o `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[0].method,
  'match',
  "'# name:John Doe, age:35' matches the regex"
);
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `match` o `notMatch`.

```js
const response = await fetch(code + '/_api/get-tests?type=unit&n=14');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(
  data.assertions[1].method,
  'notMatch',
  "'# name:Paul Smith III, age:twenty-four' does not match the regex (the age must be numeric)"
);
```
