---
id: 587d824c367417b2b2512c4c
title: "Verifica l'uguaglianza profonda con .deepEqual e .notDeepEqual"
challengeType: 2
forumTopicId: 301587
dashedName: assert-deep-equality-with--deepequal-and--notdeepequal
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`deepEqual()` verifica che due oggetti siano uguali in modo profondo.

# --instructions--

All'interno di `tests/1_unit-tests.js` sotto il test etichettato `#7` nella suite `Equality`, cambia ciascun `assert` in `assert.deepEqual` o `assert.notDeepEqual` per far passare il test (dovrebbe valutare a `true`). Non modificare gli argomenti passati agli assert.

# --hints--

Tutti i test dovrebbero passare.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.state, 'passed');
})();
```

Dovresti scegliere il metodo corretto per la prima asserzione - `deepEqual` vs. `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[0].method,
    'deepEqual',
    'The order of the keys does not matter'
  );
})();
```

Dovresti scegliere il metodo corretto per la seconda asserzione - `deepEqual` vs. `notDeepEqual`.

```js
(async () => {
  const response = await fetch(code + '/_api/get-tests?type=unit&n=6');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(
    data.assertions[1].method,
    'notDeepEqual',
    'The position of elements within an array does matter'
  );
})();
```
