---
id: 5f8884f4c46685731aabfc41
title: Esegui test funzionali usando un browser headless II
challengeType: 2
forumTopicId: 301594
dashedName: run-functional-tests-using-a-headless-browser-ii
---

# --description--

Come promemoria, questo progetto si basa sul seguente progetto iniziale clonato da <a href="https://github.com/freeCodeCamp/boilerplate-mochachai/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

# --instructions--

All'interno di `tests/2_functional-tests.js`, nel test `'Submit the surname "Vespucci" in the HTML form'` (`// #6`), automatizza quanto segue:

1. Compila il modulo con il cognome `Vespucci`
2. Premi il pulsante di invio

E all'interno della callback `pressButton`:

1. Verifica che lo stato sia OK `200`
2. Verifica che il testo all'interno dell'elemento `span#name` sia `'Amerigo'`
3. Verifica che il testo all'interno dell'elemento `span#surname` sia `'Vespucci'`
4. Verifica che l'elemento/i `span#dates` esista/esistano e che il loro numero sia `1`

Non dimenticare di rimuovere la chiamata `assert.fail()`.

# --hints--

Tutti i test dovrebbero superarsi.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.state, 'passed');
```

Dovresti verificare che la richiesta del browser headless sia andata a buon fine.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[0].method, 'browser.success');
```

Dovresti verificare che il testo all'interno dell'elemento `span#name` sia `'Amerigo'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[1].method, 'browser.text');
assert.match(data.assertions[1].args[0], /('|")span#name\1/);
assert.match(data.assertions[1].args[1], /('|")Amerigo\1/);
```

Dovresti verificare che il testo all'interno dell'elemento `span#surname` sia `'Vespucci'`.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[2].method, 'browser.text');
assert.match(data.assertions[2].args[0], /('|")span#surname\1/);
assert.match(data.assertions[2].args[1], /('|")Vespucci\1/);
```

Dovresti verificare che l'elemento `span#dates` esista e che il suo numero sia 1.

```js
const response = await fetch(code + '/_api/get-tests?type=functional&n=6');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.equal(data.assertions[3].method, 'browser.elements');
assert.match(data.assertions[3].args[0], /('|")span#dates\1/);
assert.equal(data.assertions[3].args[1], 1);
```
