---
id: 686d749d6f32f708fdf2b7ad
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In una lezione precedente, hai imparato cosa sono gli elementi `svg` e come spesso vengono usati come icone nei progetti. Nei codici reali, di solito ti affidi a librerie di icone, quindi non è necessario creare elementi `svg` da zero. Tuttavia, in questo workshop costruirai un'icona a forma di cuore per imparare gli attributi principali usati all'interno di un elemento `svg`.

Inizia creando un elemento `svg` nella pagina.

# --hints--

Dovresti avere un elemento `svg` nella pagina.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
