---
id: 686daa7ed79ceacd0b264e7d
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Dovresti annidare un elemento `path` all'interno del tuo elemento `svg` per dare forma all'immagine.

Crea un elemento `path`.

# --hints--

Dovresti avere un elemento `path` annidato all'interno del tuo elemento `svg`.

```js
const path = document.querySelector('svg path');
assert.exists(path);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    <svg>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </svg>
  </body>
</html>
```
