---
id: 66f686b8ebdb982fa8e14330
title: Crie um programa de conversão de data
challengeType: 14
dashedName: lab-date-conversion
---

# --description--

Neste laboratório, você vai programar um mecanismo de conversão de datas que converte uma determinada data para diferentes formatos. Por exemplo, a data atual `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` seria formatada nas 2 maneiras seguintes:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve criar uma variável chamada `currentDate` e atribuir a ela a data e hora atuais usando o objeto `Date`.
2. Você deve criar uma variável chamada `currentDateFormat` e atribuir a ela a string `Current Date and Time: [current date]`. Substitua `[current date]` pelo resultado da variável `currentDate`.
3. Você deve registrar o valor de `currentDateFormat` no console.
4. Você deve criar uma função chamada `formatDateMMDDYYYY` que recebe o objeto de data como parâmetro. Você pode nomear este parâmetro como quiser.
5. A função `formatDateMMDDYYYY` deve converter a data atual para o formato `month/day/year` e retornar a string `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Você deve criar uma função chamada `formatDateLong` que recebe o objeto de data como parâmetro.
7. A função `formatDateLong` deve converter a data atual para o formato `Month Day, Year` e retornar a string `Formatted Date (Month Day, Year): [formatted date]`.

**Observação:** para que os testes passem, certifique-se de usar `en-US` para o locale ao formatar as datas.

# --hints--

Você deve ter uma variável chamada `currentDate` que armazene a data e hora atuais usando o objeto `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Você deve ter uma variável chamada `currentDateFormat` que armazena a data atual no formato `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Você deve registrar o valor de `currentDateFormat` no console.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Você deve ter uma função `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

A função`formatDateMMDDYYYY` deve aceitar um único parâmetro.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Quando o objeto de data armazenar `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, a função `formatDateMMDDYYYY` deve retornar `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Você deve ter uma função `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

A função`formatDateLong` deve aceitar um único parâmetro.

```js
assert.lengthOf(formatDateLong, 1);
```

Quando o objeto de data armazenar `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, a função `formatDateLong` deve retornar `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

