---
id: 614389f601bb4f611db98563
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Chini ya kipengele chako cha `.author`, tengeneza kipengele kipya cha `div` chenye darasa la vitu la `social-icons`.

Ongeza vipengele vitano vya `a` ndani ya `div` mpya, na wape sifa za kipengele zifuatazo za `href`.

- Kipengele cha kwanza cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.facebook.com/freecodecamp`.
- Kipengele cha pili cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://twitter.com/freecodecamp`.
- Kipengele cha tatu cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://instagram.com/freecodecamp`.
- Kipengele cha nne cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.linkedin.com/school/free-code-camp`.
- Kipengele cha tano cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.youtube.com/freecodecamp`.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Kipengele chako kipya cha `div` kinapaswa kuja baada ya kipengele chako cha `.author`.

```js
assert.equal(document.querySelector('.author')?.nextElementSibling?.localName, 'div');
```

Kipengele chako kipya cha `div` kinapaswa kuwa na darasa la `social-icons`.

```js
assert.isTrue(document.querySelector('.author')?.nextElementSibling?.classList?.contains('social-icons'));
```

Kipengele chako cha `.social-icons` kinapaswa kuwa na vipengele vitano vya `a`.

```js
assert.lengthOf(document.querySelector('.social-icons')?.querySelectorAll('a'), 5);
```

Kipengele chako cha kwanza cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.facebook.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[0]?.getAttribute('href'), 'https://www.facebook.com/freecodecamp');
```

Kipengele chako cha pili cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://twitter.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[1]?.getAttribute('href'), 'https://twitter.com/freecodecamp');
```

Kipengele chako cha tatu cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://instagram.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[2]?.getAttribute('href'), 'https://instagram.com/freecodecamp');
```

Kipengele chako cha nne cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.linkedin.com/school/free-code-camp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[3]?.getAttribute('href'), 'https://www.linkedin.com/school/free-code-camp');
```

Kipengele chako cha tano cha `a` kinapaswa kuwa na sifa ya kipengele ya `href` iliyowekwa kuwa `https://www.youtube.com/freecodecamp`.

```js
assert.include(document.querySelector('.social-icons')?.querySelectorAll('a')?.[4]?.getAttribute('href'), 'https://www.youtube.com/freecodecamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
