---
id: 6143869bb45bd85e3b1926aa
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Ndani ya kipengele chako cha `.heading`, tengeneza kipengele cha `header` chenye sifa ya kipengele `class` ikiwa `hero`.

Ndani ya kipengele hicho, tengeneza kipengele cha `img` chenye sifa za kipengele `src` zikiwa `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`, `alt` zikiwa `freecodecamp logo`, na `class` zikiwa `hero-img`.

Sifa ya kipengele `loading` kwenye kipengele cha `img` inaweza kuwekwa kuwa `lazy` ili kumwambia kivinjari kutochukua rasilimali ya picha hadi itakapohitajika (yaani, wakati mtumizi anapopiga picha hadi ionekane). Faida ya ziada, vipengele vilivyo na mzigo wa polepole havitapakia hadi vipengele visivyo na mzigo wa polepole vipakie - hii inamaanisha watumizi wenye muunganisho wa intaneti polepole wanaweza kuona maudhui ya ukurasa wako bila kusubiri picha zipakie.

Mpe kipengele chako kipya cha `img` sifa ya kipengele `loading` ikiwa `lazy`.

# --hints--

Unapaswa kutengeneza kipengele cha `header`.

```js
assert.exists(document.querySelector('header'));
```

Kipengele chako cha `header` kinapaswa kuwa ndani ya kipengele chako cha `.heading`.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Kipengele chako cha `header` kinapaswa kuwa na `class` ikiwa `hero`.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

Unapaswa kuunda kipengele cha `img`.

```js
assert.exists(document.querySelector('img'));
```

Kipengele chako cha `img` kinapaswa kuwa ndani ya kipengele chako cha `header`.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Kipengele chako cha `img` kinapaswa kuwa na `src` ikiwa `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Kipengele chako cha `img` kinapaswa kuwa na `alt` ikiwa `freecodecamp logo`.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Kipengele chako cha `img` kinapaswa kuwa na sifa ya kipengele `loading` ikiwa `lazy`.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Kipengele chako cha `img` kinapaswa kuwa na `class` ikiwa `hero-img`.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
