---
id: 614387cbefeeba5f3654a291
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Baada ya kipengele chako cha `header`, tengeneza `div` na `class` kuwekwa kuwa `author`.

Ndani ya `div`, tengeneza kipengele cha `p` na `class` kuwekwa kuwa `author-name` na kipa maandishi `By freeCodeCamp`. Funika sehemu ya `freeCodeCamp` ndani ya kipengele cha `a` na `href` iwe imewekwa kuwa `https://freecodecamp.org`, na `target` iwe imewekwa kuwa `_blank`.


Chini ya hapo, ongeza kipengele cha pili cha `p` chenye darasa la vitu `publish-date` na maandishi `March 7, 2019`.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert.exists(document.querySelector('div'));
```

Kipengele chako cha `div` kinapaswa kuja baada ya kipengele chako cha `header`.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

Kipengele chako cha `div` kinapaswa kuwa na `class` iliyowekwa kuwa `author`.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

Unapaswa kutengeneza vipengele viwili vipya vya `p`.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

Vipengele vyako viwili vipya vya `p` vinapaswa kuwa ndani ya kipengele chako cha `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Kipengele chako cha kwanza kipya cha `p` kinapaswa kuwa na `class` iliyowekwa kuwa `author-name`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

Kipengele chako cha kwanza kipya cha `p` kinapaswa kuwa na maandishi `By freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

Kipengele chako cha pili kipya cha `p` kinapaswa kuwa na `class` iliyowekwa kuwa `publish-date`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

Kipengele chako cha pili kipya cha `p` kinapaswa kuwa na maandishi `March 7, 2019`.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

Unapaswa kutengeneza kipengele kipya cha `a`.

```js
assert.exists(document.querySelector('a'));
```

Kipengele chako cha `a` kinapaswa kuwa ndani ya kipengele chako kipya cha kwanza cha `p`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

Kipengele chako cha `a` kinapaswa kuwa na `href` iliyowekwa kuwa `https://freecodecamp.org`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

Kipengele chako cha `a` kinapaswa kuwa na `target` iliyowekwa kuwa `_blank`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

Kipengele chako cha `a` kinapaswa kufunika maandishi ya `freeCodeCamp`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
