---
id: 614385513d91ae5c251c2052
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utaunda ukurasa wa jarida ukitumia taarifa ya vyombo vya habari kutoka toleo la awali la mtaala wa freeCodeCamp.

Anza kwa kuongeza kipengele cha `link` kwa fonti ya mtindo wa maandishi ya `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`, kipengele cha `link` kwa fonti ya mtindo wa maandishi ya `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome, na kipengele cha `link` kwa fonti yako ya mtindo wa maandishi ya `./styles.css`.

Fonti yako ya mtindo wa maandishi itapakia fonti tatu tofauti: `Anton`, `Baskervville`, na `Raleway`.

# --hints--

Msimbo wako unapaswa kuwa na vipengele vitatu vya `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Vipengele vyako vya `link` vinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Vipengele vyako vitatu vya `link` vinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani ya `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Kimoja kati ya vipengele vyako vya kiungo kinapaswa kuwa na `href` iliyowekwa kuwa `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Kimoja kati ya vipengele vyako vya kiungo kinapaswa kuwa na `href` iliyowekwa kuwa `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Kimoja kati ya vipengele vyako vya `link` kinapaswa kuwa na `href` iliyowekwa kuwa `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
