---
id: 6165d3b702a5d92ad970b30c
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Baada ya kipengele chako cha `img`, ongeza kipengele cha `h1` chenye `class` imewekwa kuwa `hero-title` na maandishi yakiwa `OUR NEW CURRICULUM`, ikifuatiwa na kipengele cha `p` chenye `class` imewekwa kuwa `hero-subtitle` na maandishi yakiwa `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Unapaswa kuunda kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuja baada ya kipengele chako cha `img`.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

Kipengele chako cha `h1` kinapaswa kuwa na `class` imewekwa kuwa `hero-title`.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi yakiwa `OUR NEW CURRICULUM`.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

Unapaswa kuunda kipengele kipya cha `p`.

```js
assert.exists(document.querySelector('p'));
```

Kipengele chako cha `p` kinapaswa kuja baada ya kipengele chako cha `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Kipengele chako cha `p` kinapaswa kuwa na `class` imewekwa kuwa `hero-subtitle`.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Kipengele chako cha `p` kinapaswa kuwa na maandishi yakiwa `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
