---
id: 6169cd8a558aa8434e0ad7f6
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Kichwa cha HTTP cha `Referer` kina taarifa kuhusu anwani au URL ya ukurasa ambao mtumizi anaweza kuwa anautembelea kutoka. Taarifa hii inaweza kutumika katika uchambuzi kufuatilia ni watumizi wangapi kutoka ukurasa wako hutembelea freecodecamp.org, kwa mfano. Kuweka sifa ya kipengele ya `rel` kuwa `noreferrer` hutoa taarifa hii kutoka kwenye ombi la HTTP. Toa kipengele cha `a` sifa ya kipengele ya `rel` iliyowekwa kuwa `noreferrer`.

# --hints--

Kipengele chako cha `a` kinapaswa kuwa na `rel` iliyowekwa kuwa `noreferrer`.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel'), 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
--fcc-editable-region--
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
--fcc-editable-region--
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
