---
id: a5deed1811a43193f9f1c841
title: реалізувати пропуск елементів
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

У цій лабораторній роботі ви створите функцію, яка пропускає елементи в масиві, доки не знайде прийнятний на основі певної тестової функції.

Наприклад, для масиву `[1, 1, 1, 2, 1, 1, 1]` і тестової функції `function(n) {return n === 2}`, перший прийнятний елемент — це той, що має індекс 3, тому всі елементи до нього потрібно відкинути, а вихідними даними мають бути решта елементів `[2, 1, 1, 1]`.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача**

1. У вас має бути функція `dropElements`, яка приймає масив (`arr`) і функцію (`func`) як аргументи.
1. Функція `dropElements` має ітерувати по масиву і видаляти елементи, починаючи з першого, доки `func` не поверне `true` для елемента.
1. Функція `dropElements` має повертати решту елементів масиву, якщо умова виконана.
1. Якщо умова ніколи не виконується, вона має повертати порожній масив.

# --hints--

У вас має бути функція `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` має повертати `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` має повертати `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` має повертати `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` має повертати `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` має повертати `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` має повертати `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
