---
id: 655791e6cf5e03be3de73451
title: Passaggio 50
challengeType: 20
dashedName: step-50
---

# --description--

Se provi a togliere il commento dalla chiamata della tua funzione, non funzionerà. Hai un paio di bug da correggere. Il primo succede perché nel `if` annidato stai cercando di consultare un elemento che potrebbe non esistere nella tua lista `paths[node]`. Quindi, devi assicurarti che `paths[node]` non sia vuoto prima di consultare `paths[node][-1]`.

Aggiungi una condizione aggiuntiva alla tua istruzione `if` annidata per garantire che `paths[node]` non sia vuoto prima di consultare `paths[node][-1]`.

# --hints--

Dovresti aggiungere `paths[node]` come prima condizione alla tua istruzione `if` annidata. Usa l'operatore `and` per combinare le tue condizioni.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                   
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
