---
id: 655774d01daeeaa1978b99d5
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

All'inizio, tutti gli altri nodi nel grafo sono considerati a distanza infinita dal nodo sorgente, perché la distanza non è ancora stata determinata.

Crea una clausola `else` e assegna un valore infinito al nodo nel dizionario `distances`. Per farlo, usa la funzione `float()` con la stringa `'inf'` come argomento per generare un numero in virgola mobile che rappresenta l'infinito positivo.

# --hints--

Dovresti avere una clausola `else`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1else\s*:/ms));
  }
})
```

Dovresti assegnare `float('inf')` a `distances[node]` all'interno della tua nuova clausola `else`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, "else");
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*float\s*\(\s*("|')inf\1\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
        if node == start:
            distances[node] = 0
--fcc-editable-region--
```
