---
id: 65577a17564ce8a8e06c1460
title: Passaggio 35
challengeType: 20
dashedName: step-35
---

# --description--

Le comprensioni di dizionari supportano anche la sintassi condizionale `if`/`else`:

```py
{key: val_1 if condition else val_2 for key in dict}
```

Nell'esempio sopra, `dict` è il dizionario esistente. Quando `condition` valuta a `True`, `key` avrà il valore `val_1`, altrimenti `val_2`.

Usa una comprensione di dizionario per creare un dizionario basato su `graph` e assegnalo alla variabile `distances`. Assegna alla chiave il valore zero se il nodo è uguale al nodo di partenza, e infinito altrimenti. Usa `float('inf')` per ottenere quest'ultimo risultato.

# --hints--

Dovresti usare la sintassi della comprensione di dizionario per assegnare un valore alla tua variabile `distances`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;
    assert(function_body.match(/^\s{4}distances\s*=\s*\{\s*(\w+)\s*:\s*0\s+if\s+(?:\1\s*==\s*start|start\s*==\s*\1)\s+else\s+float\s*\(\s*("|')inf\2\s*\)\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    paths = {node: [] for node in graph}
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
