---
id: 65578cb031cd93b77a285db2
title: Passaggio 38
challengeType: 20
dashedName: step-38
---

# --description--

La tua funzione esplorerà tutti i nodi collegati al nodo di partenza. Calcolerà i percorsi più brevi per tutti loro. Poi, rimuoverà il nodo di partenza dai nodi non visitati.

Successivamente, verrà visitato il nodo vicino più vicino e il processo si ripeterà finché tutti i nodi saranno visitati.

Da questo momento in poi, lavorerai sul loop principale che esplora i nodi nel grafo. Per evitare problemi di esecuzione di un loop infinito durante lo sviluppo dell'algoritmo, trasforma la chiamata alla tua funzione in un commento.

# --hints--

Dovresti trasformare la chiamata alla tua funzione in un commento.

```js
({ test: () => assert.match(code, /#\s*shortest_path\s*\(\s*my_graph\s*,\s*("|')A\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}
--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
