---
id: 655790d113d14dbb727eaf41
title: Passaggio 45
challengeType: 20
dashedName: step-45
---

# --description--

Una volta impostata la distanza verso un nodo all'interno del dizionario `distances`, devi anche tenere traccia del percorso verso quel nodo. Se la distanza per il nodo nella tupla elaborata è stata aggiornata, l'ultimo elemento nel suo percorso è il nodo stesso.

All'interno della tua condizione, annida un'altra istruzione `if` che si attiva quando l'ultimo elemento di `paths[node]` è uguale a `node`. Usa `pass` per completare il corpo dell'istruzione `if`.

# --hints--

Dovresti creare un'istruzione `if` annidata che verifica se `paths[node][-1]` è uguale a `node`. Non dimenticare di usare `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)if.*:.*^\1(\s{4})if\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\2\2pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
```
