---
id: 655771d889132f9ccd341060
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Un grafo si chiama *pesato* quando i suoi archi sono associati a pesi, che rappresentano una distanza, un tempo o un altro valore quantitativo.

Nel tuo caso, questi pesi saranno le distanze tra ogni nodo, o punto nello spazio. Per rappresentare un grafo pesato puoi modificare il tuo dizionario, usando una lista di tuple per ogni valore.

Il primo elemento della tupla sarà il nodo connesso, e il secondo elemento sarà un numero intero che indica la distanza.

Modifica `my_graph['A']` in una lista di tuple, considerando che la distanza `A-B` è `3` e la distanza `A-D` è `1`.

# --hints--

`my_graph["A"]` dovrebbe essere una lista contenente le tuple `('B', 3)` e `('D', 1)`.

```js
({ test: () => assert(runPython(`
    tuples = [("B", 3), ("D", 1)]
    len(my_graph["A"]) == 2 and all(t in my_graph["A"] for t in tuples)
  `))
})
```

`my_graph` dovrebbe avere 4 chiavi chiamate `'A'`, `'B'`, `'C'` e `'D'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': ['B', 'D'],
--fcc-editable-region--    
    'B': ['A', 'C'],
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
