---
id: 65576ff7888f9e96f52a4be1
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

I dizionari archiviano dati sotto forma di coppie *chiave*-*valore*. Una chiave è separata dal valore corrispondente da due punti. E ogni coppia chiave-valore è separata dalla coppia successiva da una virgola:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Aggiungi una nuova coppia chiave-valore al tuo dizionario. Usa la stringa `'species'` come chiave e la stringa `'guinea pig'` come valore.

# --hints--

Dovresti avere una chiave `'species'` con il valore `'guinea pig'` all’interno del tuo dizionario `copper`.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Il tuo dizionario `copper` dovrebbe avere una sola coppia chiave-valore.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
