---
id: 655773b0591c5f9f4045883e
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

L'algoritmo partirà da un nodo specificato. Poi esplorerà il grafo per trovare il percorso più breve tra il nodo di partenza, o *source*, e tutti gli altri nodi.

Per questo la tua funzione ha bisogno di due parametri: `graph` e `start`. Aggiungili alla dichiarazione della tua funzione.

# --hints--

La tua funzione dovrebbe prendere `graph` e `start` come parametri, in questo ordine.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(shortest_path))
    sig == '(graph, start)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path():
    pass
--fcc-editable-region--
```
