---
id: 65579228c669fcbebffd01d5
title: Passaggio 52
challengeType: 20
dashedName: step-52
---

# --description--

L'altro bug è sottile. Quando si trova una distanza più breve per un nodo vicino, `paths[current]` viene assegnato al percorso del nodo vicino, `paths[node]`.

Questo significa che entrambe le variabili puntano alla stessa lista. Poiché le liste sono mutabili, quando aggiungi il nodo vicino al suo percorso, sia `paths[node]` che `paths[current]` vengono modificati perché sono la stessa lista. Questo porta a percorsi errati, anche se le distanze sono corrette.

Puoi correggere questo bug assegnando una copia di `paths[current]` al percorso del nodo vicino. Per farlo puoi usare la sintassi di slicing:

```py
my_list[:]
```

Dove `my_list` è la lista che vuoi copiare. Modifica l'assegnazione `paths[node] = paths[current]` esistente all'interno del tuo blocco `if` facendo lo slicing di `paths[current]`.

# --hints--

Dovresti assegnare `paths[current][:]` al percorso del nodo vicino.

```js
({ test: () => assert.match(code, /^(\s*)if\s+paths\s*\[\s*node\s*\]\s+and\s+paths\s*\[\s*node\s*\]\s*\[\s*-\s*1\s*\]\s*==\s*node\s*:\s*^\1\s{4}paths\s*\[\s*node\s*\]\s*=\s*paths\s*\[\s*current\s*\]\s*\[\s*::?\s*\]/ms) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    while unvisited:
        current = min(unvisited, key=distances.get)
        for node, distance in graph[current]:
            if distance + distances[current] < distances[node]:
                distances[node] = distance + distances[current]
--fcc-editable-region--                
                if paths[node] and paths[node][-1] == node:
                    paths[node] = paths[current]
--fcc-editable-region--                    
                else:
                    paths[node].extend(paths[current])
                paths[node].append(node)
        unvisited.remove(current)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
shortest_path(my_graph, 'A')

```
