---
id: 655773f8b8b5db9fc6d0ae76
title: Passaggio 25
challengeType: 20
dashedName: step-25
---

# --description--

Per tenere traccia dei nodi visitati, ti serve una lista di tutti i nodi nel grafo. Una volta che un nodo è visitato, verrà rimosso da quella lista.

Ora, sostituisci la parola chiave `pass` con una variabile chiamata `unvisited` e assegnale una lista vuota.

# --hints--

Dovresti avere una variabile chiamata `unvisited` all’interno della funzione `shortest_path`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").has_variable("unvisited")`)) })
```

Dovresti assegnare una lista vuota alla tua variabile `unvisited`. Ricorda di eliminare `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_body().is_equivalent("unvisited = []")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    pass
--fcc-editable-region--
```
