---
id: 65577791ad8c26a7705e2919
title: Passaggio 34
challengeType: 20
dashedName: step-34
---

# --description--

Con una comprensione del dizionario, puoi creare un dizionario partendo da un dizionario esistente:

```py
{key: val for key in dict}
```

Nell'esempio sopra, `val` è il valore che `key` avrà nel nuovo dizionario, e `dict` è il dizionario esistente.

Vuoi tenere traccia dei percorsi tra il nodo di partenza e ogni altro nodo.

Dopo la variabile `distances`, crea una variabile `paths` e assegnale un dizionario con tutte le chiavi di `graph`. Assegna a ogni chiave una lista vuota e usa una comprensione del dizionario per buildare il tuo dizionario.

# --hints--

Dovresti avere una variabile chiamata `paths`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^\s{4}paths\s*\=/m));
  }
})
```

La tua variabile `paths` dovrebbe usare la sintassi della comprensione del dizionario per assegnare una lista vuota a ogni nodo in graph.

```js
({ test: () =>  {
    const {function_body} = __helpers.python.getDef(code, "shortest_path");       
    assert(function_body.match(/^\s{4}paths\s*\=\s*\{\s*(\w+)\s*:\s*\[\s*\]\s+for\s+\1\s+in\s+graph\s*\}/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
