---
id: 65577236b056379d5dbc7000
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Ora inizierai a sviluppare l'algoritmo per calcolare il percorso più breve tra ogni nodo nel tuo nuovo grafo.

Dichiara una funzione vuota chiamata `shortest_path`. Usa la parola chiave `pass` per completare il corpo della funzione.

# --hints--

Dovresti definire una funzione chiamata `shortest_path`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(shortest_path)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}


--fcc-editable-region--
```
