---
id: 65578d0f6c78a0b868a43b9c
title: Passaggio 40
challengeType: 20
dashedName: step-40
---

# --description--

All'interno del loop `while`, la prima cosa da fare è definire il nodo corrente da visitare. Per questo puoi usare la funzione `min()`. Restituisce l'elemento più piccolo dall'iterabile passato come argomento.

Elimina `pass`, poi crea una variabile chiamata `current` e assegnale `min(unvisited)`.

# --hints--

Dovresti creare una variabile `current` nel tuo loop `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

Dovresti assegnare `min(unvisited)` alla tua variabile `current`. Ricorda di eliminare `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
