---
id: 655774955b097ea14897db12
title: Passaggio 28
challengeType: 20
dashedName: step-28
---

# --description--

La distanza dal nodo di partenza è zero, perché l’algoritmo inizia la sua valutazione proprio da lì.

Dopo aver aggiunto `node` a `unvisited` nel tuo loop, crea un’istruzione `if` che si attivi se il nodo è uguale al nodo di partenza. Poi assegna `0` a quel nodo all’interno del dizionario `distances`.

# --hints--

Dovresti creare un’istruzione `if` che venga eseguita quando `node` è uguale a `start`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+node\s+in\s+graph\s*/);        
    assert(block_body.match(/unvisited\.append\s*\(\s*node\s*\)\s*^\s+if\s+(node\s*==\s*start|start\s*==\s*node)\s*:/m));
  }
})
```

All’interno della tua nuova istruzione `if` dovresti assegnare `0` al nodo nel dizionario `distances`.

```js
({ test: () =>  {
   const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+(node\s*==\s*start|start\s*==\s*node)\s*/m);  
    assert(block_body.match(/^\s+distances\s*\[\s*node\s*\]\s*=\s*0/m));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    for node in graph:
        unvisited.append(node)
--fcc-editable-region--
```
