---
id: 65578d4fc3afc3b8f554c882
title: Passaggio 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` accetta anche un argomento keyword-only. Passando una funzione come argomento aggiuntivo a `min()`, puoi modificare il modo in cui vengono confrontati gli elementi della lista.

Il risultato della riga che hai appena scritto nel passaggio precedente è il nodo che viene per primo in ordine alfabetico. Invece vuoi selezionare il nodo non visitato che ha la distanza minore dal nodo di partenza.

Passa `key=distances.get` come secondo argomento alla tua chiamata `min()`. In questo modo, il confronto avverrà in base al valore che ogni elemento della lista `unvisited` ha all’interno del dizionario `distances`.

# --hints--

Dovresti passare `key=distances.get` come secondo argomento alla tua chiamata `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
