---
id: 6557709b0aee699a6a00528c
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ad esempio, un grafo può essere usato per rappresentare due punti nello spazio, `A` e `B`, collegati da un percorso. Un grafo di questo tipo sarà composto da due nodi collegati da un arco.

Sostituisci la chiave `'species'` esistente con le stringhe `'A'`. Poi, sostituisci il valore corrispondente con la stringa `'B'` per rappresentare la connessione tra i nodi `'A'` e `'B'`.

# --hints--

Il tuo dizionario dovrebbe avere una chiave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

La tua chiave `'A'` dovrebbe avere `'B'` come valore.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Il tuo dizionario dovrebbe avere due chiavi.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
