---
id: 65577739f57ecca6c39bb4e9
title: Passaggio 33
challengeType: 20
dashedName: step-33
---

# --description--

Il costruttore del tipo `list()` ti permette di creare una lista da un iterabile.

Modifica l'assegnazione della tua variabile `unvisited` per usare `list()` e passa `graph` come iterabile.

# --hints--

Dovresti usare `list()` per generare una lista dal dizionario `graph`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/list\s*\(\s*graph\s*\)/));
  }
})
```

Dovresti assegnare `list(graph)` a `unvisited`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/unvisited\s*=\s*list\s*\(\s*graph\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = []
    distances = {}
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}')
    
shortest_path(my_graph, 'A')
--fcc-editable-region--
```
