---
id: 65578cee7f2cb8b80127cce2
title: Passaggio 39
challengeType: 20
dashedName: step-39
---

# --description--

Prima della chiamata `print`, crea un loop `while` che venga eseguito finché `unvisited` non è vuoto. Usa la parola chiave `pass` per riempire il corpo del loop.

# --hints--

Dovresti avere un loop `while` che si esegue mentre `unvisited` non è vuoto. Non dimenticare la parola chiave `pass`.

```js
({ test: () =>  {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "shortest_path");     
    assert(function_body.match(/^\s+while\s+(unvisited|unvisited\s*!=\s*\[\s*\]|len\s*\(\s*unvisited\s*\)\s*(>|!=)\s*0)\s*:/m));
    const {block_body} = __helpers.python.getBlock(commentless_code, /while\s+(unvisited|unvisited\s*!=\s*\[\s*\]|len\s*\(\s*unvisited\s*\)\s*(>|!=)\s*0)\s*/);
    assert(block_body.match(/\s+pass/))
    
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

--fcc-editable-region--
def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')
--fcc-editable-region--
```
