---
id: 6557712d77ce2d9bd7e63afd
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Aggiungi un altro nodo collegato a `B` al tuo grafo e chiamalo `C`.

Modifica il dizionario esistente per rappresentare questa disposizione: aggiungi un’altra chiave `'C'` a `my_graph` e assegna a essa il valore della stringa `'B'`.

Inoltre, cambia il valore della chiave `'B'` esistente in una lista `['A', 'C']` per rappresentare le connessioni multiple del tuo nodo `'B'`.

# --hints--

Il tuo dizionario dovrebbe avere 3 chiavi — `'A'`, `'B'` e `'C'`.

```js
({ test: () => assert(runPython(`
    key_list = ["A", "B", "C"]
    len(my_graph) == 3 and all(key in my_graph for key in key_list)
  `))
})
```

Il valore di `my_graph['A']` dovrebbe essere la stringa `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

`my_graph['B']` dovrebbe essere una lista.

```js
({ test: () => assert(runPython(`
    type(my_graph["B"]) is list
  `))
})
```

Il valore di `my_graph['B']` dovrebbe essere una lista contenente `'A'` e `'C'`.

```js
({ test: () => assert(runPython(`
    len(my_graph["B"]) == 2 and "A" in my_graph["B"] and "C" in my_graph["B"]
  `))
})
```

Il valore di `my_graph['C']` dovrebbe essere la stringa `'B'`.

```js
({ test: () => assert(runPython(`
    my_graph["C"] == "B"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'B': 'A'
}
--fcc-editable-region--
```
