---
id: 65578f895f2a65ba7a916804
title: Passaggio 42
challengeType: 20
dashedName: step-42
---

# --description--

Dopo l'assegnazione della variabile `current`, crea un loop `for` per iterare sulle tuple nella lista `graph[current]`. Ti serviranno due variabili per iterare. Ricorda di usare `pass` per riempire il corpo del loop.

# --hints--

Dovresti creare un loop `for` per iterare sugli elementi tuple nella lista `graph[current]`. Usa due variabili per iterare e non dimenticare di aggiungere la parola chiave `pass`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s*)while\s+unvisited\s*:.*^\1\1for\s+\w+\s*,\s*\w+\s+in\s+graph\s*\[\s*current\s*\]\s*:\s*^\1\1\1pass/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited, key=distances.get)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')
    
#shortest_path(my_graph, 'A')

```
