---
id: 6566195b0a021bb660b2b4b1
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Ora modifica `my_graph['B']` in una lista di tuple, dove il primo elemento della tupla è il nodo connesso e il secondo elemento è la distanza. La distanza `B-C` è `4`.

# --hints--

`my_graph['B']` dovrebbe essere una lista contenente le tuple `('A', 3)` e `('C', 4)`.

```js
({ test: () => assert(runPython(`
    tuples = [("A", 3), ("C", 4)]
    len(my_graph["B"]) == 2 and all(t in my_graph["B"] for t in tuples)
  `))
})
```

`my_graph` dovrebbe avere 4 chiavi chiamate `'A'`, `'B'`, `'C'` e `'D'`.

```js
({ test: () => assert(runPython(`    
    key_list = ["A", "B", "C", "D"]
    len(my_graph) == 4 and all(key in my_graph for key in key_list)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': ['A', 'C'],
--fcc-editable-region--    
    'C': ['B', 'D'],
    'D': ['A', 'C']
}

```
