---
id: 65789506b30453080f77470c
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Finora hai già incontrato diversi tipi di dati:

- Tipi di dati immutabili, come numeri interi, stringhe, tuple e booleani.
- Tipi di dati mutabili, come liste e dizionari.

Un dizionario è identificato da una coppia di parentesi graffe, `{}`.

Inizia creando una variabile chiamata `copper` e assegnale un dizionario vuoto usando una coppia di parentesi graffe, allo stesso modo in cui creeresti una lista vuota con una coppia di parentesi quadre.

# --hints--

Dovresti avere una variabile chiamata `copper`.

```js
({ test: () => assert(__userGlobals.has("copper")) })
```

La tua variabile `copper` dovrebbe avere il valore di un dizionario vuoto. Usa una coppia di parentesi graffe per questo.

```js
({ test: () => assert(runPython(`
    copper == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
