---
id: 6578b57361f2f132a02e2a18
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Puoi consultare i dati archiviati in un dizionario tramite le sue chiavi:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['name'] # 'Michael'
```

Dopo il tuo dizionario, segui l'esempio sopra per consultare la chiave `'species'` di `copper` e stampare il risultato.

# --hints--

Non dovresti modificare l'assegnazione del tuo dizionario.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Dovresti usare `copper['species']` per consultare il valore della chiave `'species'`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')species\1\s*\]/) })
```

Dovresti chiamare `print()` passando `copper['species']` come argomento.

```js
({ test: () => assert.match(code, /^print\s*\(\s*copper\s*\[\s*("|')species\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}

--fcc-editable-region--
```
