---
id: 6578b13757611e2825beb8a5
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Le chiavi devono essere uniche all'interno di un dizionario e possono essere solo tipi di dati immutabili. Questo significa che non puoi usare una lista o un altro dizionario come chiavi.

Aggiungi un'altra chiave `'age'` al tuo dizionario e assegna come valore il numero intero `2`.

# --hints--

Dovresti avere una chiave `'age'` con il valore `2` subito dopo `'species': 'guinea pig'`, all'interno del tuo dizionario `copper`. Non dimenticare la virgola.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig", "age": 2}
  `))
})
```

Dovresti avere due coppie chiave-valore nel tuo dizionario `copper`.

```js
({ test: () => assert(runPython(`
    len(copper) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {'species': 'guinea pig'}
--fcc-editable-region--
```
