---
id: 6579717f0920131304286804
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Per aggiungere una nuova coppia chiave-valore dopo aver dichiarato un dizionario, puoi indicare la chiave nello stesso modo in cui consulteresti una chiave esistente e assegnare il valore della nuova chiave usando l'operatore di assegnazione:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

my_dict['country'] = 'Canada'
```

Elimina la tua chiamata `print()`. Poi, dopo la dichiarazione `copper`, aggiungi la chiave `'food'` al tuo dizionario e assegna il suo valore a `'hay'`.

# --hints--

Non dovresti avere `print(copper['age'])` nel tuo codice.

```js
({ test: () => assert.notMatch(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

Dovresti aggiungere la chiave `'food'` a `copper` dopo aver dichiarato il dizionario.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')food\1\s*\]/) })
```

Dovresti assegnare `copper['food']` a `'hay'` dopo aver dichiarato il dizionario.

```js
({ test: () => assert.match(code, /^copper\s*\[\s*("|')food\1\s*\]\s*=\s*("|')hay\2/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['age'])
--fcc-editable-region--
```
