---
id: 6579dd49fa8a8e1fd06b85a9
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

I grafi sono strutture dati che rappresentano relazioni tra coppie di elementi.  
Questi elementi, chiamati *nodi*, possono essere oggetti reali, entità, punti nello spazio o altro. Le connessioni tra i nodi si chiamano *archi*.  

Ecco una rappresentazione visiva di un grafo:

<img alt="a weighted graph with 6 nodes" src="https://cdn.freecodecamp.org/curriculum/python/graph1-example.png" style="background-color: white; height: 350px; width: auto; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">


Rinomina il dizionario `copper` in `my_graph`. Questo rappresenterà il grafo per testare il tuo algoritmo.

# --hints--

Dovresti rinominare il tuo dizionario `copper` in `my_graph`.

```js
({ test: () => assert(__userGlobals.has("my_graph")) })
```

La tua variabile `my_graph` dovrebbe essere un dizionario.

```js
({ test: () => assert(runPython(`
    type(my_graph) is dict 
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
```
