---
id: 6579cbab9825b8170974c69a
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Puoi eliminare una coppia chiave-valore da un dizionario usando la parola chiave `del`:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}

del my_dict['occupation']
```

Subito prima del tuo loop `for`, usa la parola chiave `del` per eliminare la chiave `'age'` e il suo valore da `copper`.

# --hints--

Dovresti usare la parola chiave `del` per eliminare `copper['age']` prima del loop `for`.

```js
({ test: () => assert.match(code, /^del\s+copper\s*\[\s*("|')age\1\s*\].*^for\s*/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
