---
id: 65796fac81f983127558f3f4
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora, modifica la tua chiamata `print()` esistente per stampare il valore della chiave `age`.


# --hints--

Dovresti usare `copper['age']` per consultare il valore della chiave `age`.

```js
({ test: () => assert.match(code, /copper\s*\[\s*("|')age\1\s*\]/) })
```

Dovresti chiamare `print()` passando `copper['age']` come argomento.

```js
({ test: () => assert.match(code, /^print\(\s*copper\s*\[\s*("|')age\1\s*\]\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
print(copper['species'])
--fcc-editable-region--
```
