---
id: 6579ca0923cfa7162089d2f0
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

La stessa sintassi può essere usata per cambiare il valore di una chiave esistente.

Subito prima della chiamata `print()`, consulta la chiave `'species'` e riassegna il suo valore a `'Cavia porcellus'`.

# --hints--

Dovresti riassegnare `copper['species']` alla stringa `'Cavia porcellus'` prima della chiamata `print()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /^copper\s*\[\s*("|')species\1\s*\]\s*=\s*("|')Cavia porcellus\2\s+^print\s*\(\s*copper\s*\)/m)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'

print(copper)
--fcc-editable-region--
```
