---
id: 6579cd5f6dd62c189e53ddbb
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Ora che hai rivisto gli aspetti base dei dizionari, puoi procedere a costruire l'algoritmo del percorso più breve.

Elimina tutte le righe di codice dopo la dichiarazione del dizionario `copper`.

# --hints--

Dovresti eliminare tutte le righe di codice dopo la dichiarazione del tuo dizionario.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse( /copper\s*\[.*?\]|del|for|print\s*\(.*?\)/ms.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
copper = {
    'species': 'guinea pig',
    'age': 2
}
--fcc-editable-region--
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'
del copper['age']

for i, j in copper.items():
    print(i, j)
--fcc-editable-region--
```
