---
id: 65affe602a2a828b0382a058
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Sostituisci la chiave `'age'` con la stringa `'B'` e imposta il suo valore sulla stringa `'A'` per rappresentare la connessione tra i nodi in entrambe le direzioni.

# --hints--

Il tuo dizionario dovrebbe avere una chiave `'B'`.

```js
({ test: () => assert(runPython(`
    "B" in my_graph
  `))
})
```

La tua chiave `'B'` dovrebbe avere `'A'` come valore.

```js
({ test: () => assert(runPython(`
    my_graph["B"] == "A"
  `))
})
```

Il tuo dizionario dovrebbe avere una chiave `'A'`.

```js
({ test: () => assert(runPython(`
    "A" in my_graph
  `))
})
```

La tua chiave `'A'` dovrebbe avere `'B'` come valore.

```js
({ test: () => assert(runPython(`
    my_graph["A"] == "B"
  `))
})
```

Il tuo dizionario dovrebbe avere due chiavi.

```js
({ test: () => assert(runPython(`
    len(my_graph) == 2
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_graph = {
    'A': 'B',
    'age': 2
}
--fcc-editable-region--
```
