---
id: 65b7cd2b7bd9a684ccf1dc16
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Per iterare sulle chiavi di un dizionario, puoi semplicemente mettere il dizionario in un loop `for`. Il codice qui sotto stamperebbe ogni chiave nel dizionario `dict`:

```py
for i in dict:
   print(i)
```

Sostituisci la chiamata `print()` con un loop `for` che itera su `copper` e stampa ogni chiave.

# --hints--

Non dovresti avere `print(copper)` nel tuo codice.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("print(copper)")`)) })
```

Dovresti creare un loop `for` per iterare sul dizionario `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper")`)) })
```

Dovresti stampare ogni chiave di `copper` all'interno del tuo loop `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print({var})")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

print(copper)
--fcc-editable-region--
```
