---
id: 65b7d3319ebec69b983fb91b
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Come puoi vedere dall'output, `.items()` crea una struttura dati che archivia ogni coppia chiave-valore in una tupla distinta. Per iterare sugli elementi di quelle tuple puoi aggiungere una seconda variabile di loop:

```py
for i, j in dict.items():
    print(i, j)
```

Modifica il tuo loop `for` per usare due variabili di loop e stampare entrambe all'interno del corpo del loop.

# --hints--

Dovresti modificare il tuo loop `for` per usare due variabili.

```js
({ test: () => assert(runPython(`
import ast
var = _Node(_code).find_for_loops()[0].find_for_vars()
isinstance(var.tree, ast.Tuple)
`)) })
```

Dovresti stampare entrambe le variabili di loop all'interno del tuo loop `for`.

```js
({ test: () => assert(runPython(`
var = _Node(_code).find_for_loops()[0].find_for_vars()
_Node(_code).find_for_loops()[0].find_bodies()[0].is_equivalent(f"print{var}")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.items():
    print(i)
--fcc-editable-region--
```
