---
id: 65b7cf140d34058d7ea8935f
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Se vuoi iterare sui valori delle chiavi del dizionario, un modo è usare il metodo `.values()`.

Modifica il tuo loop `for` per iterare su `copper.values()` invece di `copper` e osserva l'output.

# --hints--

Dovresti modificare il tuo loop `for` per iterare su `copper.values()` invece di `copper`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.values()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper:
    print(i)
--fcc-editable-region--
```
