---
id: 65b7d25992879698180e6a71
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Infine, se vuoi poter iterare su coppie chiave-valore, puoi usare il metodo `.items()`.

Modifica il tuo loop `for` per iterare su `copper.items()` invece di `copper.values()`.

# --hints--

Dovresti modificare il tuo loop `for` per iterare su `copper.items()` invece di `copper.values()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_for_loops()[0].find_for_iter().is_equivalent("copper.items()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {
    'species': 'guinea pig',
    'age': 2
}
copper['food'] = 'hay'
copper['species'] = 'Cavia porcellus'

for i in copper.values():
    print(i)
--fcc-editable-region--
```
