---
id: 587d7b8a367417b2b2512b4c
title: rest要素による分割代入
challengeType: 1
forumTopicId: 301218
dashedName: >-
  use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

配列の分割代入で、残りの要素を別の配列にまとめたい場合があります。

結果は以下の`Array.prototype.slice()`のようになります。

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

コンソールには`1, 2`と`[3, 4, 5, 7]`の値が表示されます。

変数`a`と`b`は配列の最初と2番目の値を受け取ります。その後、rest構文があるため、`arr`は残りの値を配列の形で受け取ります。rest要素はリストの最後の変数としてのみ正しく機能します。つまり、元の配列の最後の要素を除いたサブ配列をrest構文で受け取ることはできません。

# --instructions--

rest構文を使った分割代入で`Array.prototype.slice()`の動作を再現してください。`removeFirstTwo()`は元の配列`list`の最初の2つの要素を除いたサブ配列を返す必要があります。

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])`は`[3, 4, 5]`にする必要があります。

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

`removeFirstTwo()`は`list`を変更してはいけません。

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

`Array.slice()`は使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

rest構文を使う必要があります。

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
