---
id: 587d7b88367417b2b2512b47
title: 関数のパラメータで Rest パラメータを使う
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

より柔軟な関数を作成できるように、ES6では関数のパラメータに<dfn>rest パラメータ</dfn>が導入されました。rest パラメータを使うと、可変個の引数を受け取る関数を作成できます。これらの引数は配列に格納され、関数内から後でアクセスできます。

このコードを見てみましょう：

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

コンソールには文字列 `You have passed 3 arguments.` と `You have passed 4 arguments.` が表示されます。

rest パラメータを使うことで `arguments` オブジェクトを使う必要がなくなり、関数 `howMany` に渡されたパラメータの配列に対して配列メソッドを使うことができます。

# --instructions--

関数 `sum` を rest パラメータを使うように修正し、関数 `sum` が任意の数の引数を受け取り、それらの合計を返せるようにしてください。

# --hints--

`sum(0,1,2)` の結果は 3 になるはずです

```js
assert(sum(0, 1, 2) === 3);
```

`sum(1,2,3,4)` の結果は 10 になるはずです

```js
assert(sum(1, 2, 3, 4) === 10);
```

`sum(5)` の結果は 5 になるはずです

```js
assert(sum(5) === 5);
```

`sum()` の結果は 0 になるはずです

```js
assert(sum() === 0);
```

`sum` は rest パラメータ構文（`...`）を `args` パラメータに使うアロー関数である必要があります。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
