---
id: 587d7b8a367417b2b2512b4d
title: 関数のパラメータとしてオブジェクトを渡すために分割代入を使う
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

場合によっては、関数の引数自体でオブジェクトを分割代入できます。

以下のコードを見てください：

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

これは関数に渡されたオブジェクトを効果的に分割代入しています。これをインラインで行うこともできます：

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

`profileData` が上記の関数に渡されると、関数のパラメータから値が分割代入されて関数内で使われます。

# --instructions--

関数 `half` の引数内で分割代入を使い、関数内に `max` と `min` のみを渡してください。

# --hints--

`stats` は `object` であるべきです。

```js
assert(typeof stats === 'object');
```

`half(stats)` は `28.015` であるべきです。

```js
assert(half(stats) === 28.015);
```

分割代入を使うべきです。

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

分割代入されたパラメータを使うべきです。

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
