---
id: 587d7b87367417b2b2512b43
title: アロー関数を使って簡潔な無名関数を書く
challengeType: 1
forumTopicId: 301211
dashedName: use-arrow-functions-to-write-concise-anonymous-functions
---

# --description--

JavaScriptでは、特に関数を別の関数の引数として渡す場合、関数に名前を付ける必要がないことがよくあります。その代わりに、インライン関数を作成します。これらの関数は他で再利用しないため、名前を付ける必要がありません。

これを実現するために、次の構文をよく使います。

```js
const myFunc = function() {
  const myVar = "value";
  return myVar;
}
```

ES6では、このように無名関数を書く必要がない構文糖が提供されています。代わりに、**アロー関数構文**を使うことができます。

```js
const myFunc = () => {
  const myVar = "value";
  return myVar;
}
```

関数本体がなく、戻り値だけの場合、アロー関数構文ではキーワード`return`やコードを囲む括弧を省略できます。これにより、小さな関数を1行の文に簡潔にまとめられます。

```js
const myFunc = () => "value";
```

このコードはデフォルトで文字列`value`を返します。

# --instructions--

`magic`に代入されている、`new Date()`を返す関数をアロー関数構文を使って書き直してください。また、キーワード`var`を使って定義されているものがないようにしてください。

# --hints--

`var`キーワードを置き換える必要があります。

```js
assert.notMatch(code, /var/g)
```

`magic`は定数変数（`const`を使う）にしてください。

```js
assert.match(code, /const\s+magic/g)
```

`magic`は`function`であるべきです。

```js
assert(typeof magic === 'function');
```

`magic()`は正しい日付を返す必要があります。

```js
assert(magic().setHours(0, 0, 0, 0) === new Date().setHours(0, 0, 0, 0));
```

`function`キーワードは使わないでください。

```js
assert.notMatch(code, /function/g)
```

# --seed--

## --seed-contents--

```js
var magic = function() {
  return new Date();
};
```

# --solutions--

```js
const magic = () => {
  return new Date();
};
```
