---
id: 587d7b8c367417b2b2512b54
title: オブジェクトへのアクセスを制御するためにゲッターとセッターを使う
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

オブジェクトから値を取得したり、オブジェクト内のプロパティの値を設定したりできます。

これらは古典的に<dfn>ゲッター</dfn>と<dfn>セッター</dfn>と呼ばれます。

ゲッター関数は、ユーザーがオブジェクトのプライベート変数に直接アクセスせずに、その値を単に返す（取得する）ことを目的としています。

セッター関数は、渡された値に基づいてオブジェクトのプライベート変数の値を変更（設定）することを目的としています。この変更は計算を含む場合や、以前の値を完全に上書きする場合もあります。

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

コンソールには文字列 `anonymous` と `newAuthor` が表示されます。

ゲッターとセッターを呼び出す際の構文に注目してください。関数のようには見えません。ゲッターとセッターは内部の実装の詳細を隠すために重要です。

**注意:** プライベート変数の名前の前にアンダースコア（`_`）を付けるのが慣例です。ただし、この慣例自体が変数をプライベートにするわけではありません。

# --instructions--

`class` キーワードを使って `Thermostat` クラスを作成してください。`constructor` は華氏温度を受け取ります。

クラス内に、摂氏温度を取得するための `getter` と、摂氏温度を受け取る `setter` を作成してください。

`C = 5/9 * (F - 32)` と `F = C * 9.0 / 5 + 32` を覚えておいてください。ここで `F` は華氏温度の値、`C` は同じ温度の摂氏温度の値です。

**注意:** 実装時には、クラス内で温度をどちらか一方の単位（華氏または摂氏）で管理します。

これがゲッターとセッターの力です。別のユーザー向けにAPIを作成しており、どちらの単位で管理していても正しい結果を取得できます。

つまり、ユーザーから実装の詳細を抽象化しているのです。

# --hints--

`Thermostat` は `class` で、`constructor` メソッドが定義されているべきです。

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

`class` キーワードを使うべきです。

```js
assert.match(code, /class/);
```

`Thermostat` はインスタンス化できるべきです。

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

華氏の値でインスタンス化されたとき、`Thermostat` は正しい `temperature` を設定するべきです。

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

`getter` を定義するべきです。

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

`setter` を定義するべきです。

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

摂氏の値で `setter` を呼び出すと、`temperature` が設定されるべきです。

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
