---
id: 587d7b8b367417b2b2512b50
title: ES6で簡潔な宣言的関数を書く
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

ES5でオブジェクト内に関数を定義する場合、次のようにキーワード`function`を使う必要があります。

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

ES6では、オブジェクト内で関数を定義するときに`function`キーワードとコロンを完全に省略できます。以下はこの構文の例です。

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

オブジェクト`setGear`内の関数`bicycle`を、上記の省略構文を使うようにリファクタリングしてください。

# --hints--

従来の関数式は使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear`は宣言的関数であるべきです。

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)`は`gear`の値を48に変更する必要があります。

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
