---
id: 587d7b89367417b2b2512b4a
title: ネストされたオブジェクトから変数を代入するために分割代入を使う
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

前の2つのレッスンの原則を使って、ネストされたオブジェクトから値を分割代入できます。

前の例に似たオブジェクトを使うと：

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

オブジェクトのプロパティの値を抽出し、同じ名前の変数に代入する方法は次の通りです：

```js
const { johnDoe: { age, email }} = user;
```

そして、オブジェクトのプロパティの値を異なる名前の変数に代入する方法は次の通りです：

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

2つの代入を同等の分割代入に置き換えてください。`lowToday`オブジェクトから`highToday`と`today.low`の値を`today.high`と`LOCAL_FORECAST`の変数に代入する必要があります。

# --hints--

ES5の代入構文は削除してください。

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

`lowToday`変数を作るために分割代入を使うべきです。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`highToday`変数を作るために分割代入を使うべきです。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday`は`64`と等しく、`highToday`は`77`と等しくなるべきです。

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
