---
id: 587d7b8c367417b2b2512b57
title: ファイルからすべてをインポートするために * を使う
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

ファイルがあり、その内容すべてを現在のファイルにインポートしたいとします。これは `import * as` 構文で行えます。以下は、同じディレクトリにある `math_functions.js` という名前のファイルの内容をインポートする例です。

```js
import * as myMathModule from "./math_functions.js";
```

上記の `import` 文は `myMathModule` というオブジェクトを作成します。これは単なる変数名で、好きな名前を付けられます。このオブジェクトには `math_functions.js` からのすべてのエクスポートが含まれるため、他のオブジェクトのプロパティのように関数にアクセスできます。インポートした `add` と `subtract` 関数を使う方法は以下の通りです。

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

このファイルのコードは、現在のファイルと同じディレクトリにある `string_functions.js` ファイルの内容を必要としています。`import * as` 構文を使って、そのファイルからすべてを `stringFunctions` というオブジェクトにインポートしてください。

# --hints--

コードは `import * as` 構文を正しく使う必要があります。

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
