---
id: 587d7b8c367417b2b2512b56
title: コードブロックを共有するために export を使う
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

`math_functions.js` というファイルがあり、数学的な操作に関連するいくつかの関数が含まれていると想像してください。そのうちの1つは `add` という変数に格納されていて、2つの数字を受け取り、その合計を返します。この関数を複数の異なる JavaScript ファイルで使いたい場合、まずそれを `export` する必要があります。

```js
export const add = (x, y) => {
  return x + y;
}
```

上記は単一の関数をエクスポートする一般的な方法ですが、次のようにして同じことを実現できます。

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

変数や関数をエクスポートすると、別のファイルでインポートしてコードを書き直すことなく使えます。エクスポートしたいものごとに最初の例を繰り返すか、2番目の例の export 文にすべてまとめて書くことで複数のものをエクスポートできます。

```js
export { add, subtract };
```

# --instructions--

エディターには文字列に関する関数が2つあります。どちらもお好きな方法でエクスポートしてください。

# --hints--

`uppercaseString` を正しくエクスポートしてください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

`lowercaseString` を正しくエクスポートしてください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
