---
id: 587d7b8a367417b2b2512b4f
title: オブジェクトプロパティの省略記法を使って簡潔なオブジェクトリテラル宣言を書く
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6はオブジェクトリテラルを簡単に定義できる便利な機能を追加しました。

次のコードを考えてみましょう：

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition`は2つのプロパティを含むオブジェクトを返すシンプルな関数です。ES6は`x: x`を繰り返し書く冗長さをなくす構文糖衣を提供します。`x`を一度書くだけで、内部的には`x: x`（または同等のもの）に変換されます。以下は上記の関数をこの新しい構文を使って書き直したものです：

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

オブジェクトリテラルでオブジェクトプロパティの省略記法を使い、`name`、`age`、`gender`のプロパティを持つオブジェクトを作成して返してください。

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")`は`{name: "Zodiac Hasbro", age: 56, gender: "male"}`を返すべきです。

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

コード内で`key:value`を使わないでください。

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
