---
id: 587d7b88367417b2b2512b46
title: 関数のデフォルトパラメータを設定する
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

より柔軟な関数を作成できるように、ES6では関数の<dfn>デフォルトパラメータ</dfn>が導入されました。

このコードを見てみましょう：

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

コンソールには文字列 `Hello John` と `Hello Anonymous` が表示されます。

引数が指定されていない（undefinedの場合）ときにデフォルトパラメータが適用されます。上の例のように、パラメータ `name` に値が渡されない場合、デフォルト値 `Anonymous` が割り当てられます。デフォルト値は必要なだけ複数のパラメータに設定できます。

# --instructions--

関数 `increment` を修正して、`number` が指定されていない場合に `value` に 1 を加えるようにデフォルトパラメータを追加してください。

# --hints--

`increment(5, 2)` の結果は `7` になるはずです。

```js
assert(increment(5, 2) === 7);
```

`increment(5)` の結果は `6` になるはずです。

```js
assert(increment(5) === 6);
```

`1` にはデフォルトパラメータ値 `value` を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
