---
id: 587d7b8c367417b2b2512b58
title: export default を使ってエクスポートのフォールバックを作成する
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

`export` レッスンでは、<dfn>名前付きエクスポート</dfn>と呼ばれる構文について学びました。これにより、複数の関数や変数を他のファイルで使えるようにできます。

もう一つ知っておくべき `export` 構文があり、これは <dfn>export default</dfn> と呼ばれます。通常、ファイルから1つの値だけをエクスポートする場合にこの構文を使います。また、ファイルやモジュールのフォールバック値を作成するためにも使われます。

以下は `export default` を使った例です：

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

最初は名前付き関数で、2つ目は無名関数です。

`export default` はモジュールやファイルのフォールバック値を宣言するために使われるので、各モジュールやファイルでデフォルトエクスポートは1つだけにできます。さらに、`export default` は `var`、`let`、または `const` と一緒に使うことはできません。

# --instructions--

次の関数をモジュールのフォールバック値にしてください。必要なコードを追加してください。

# --hints--

コードは `export` のフォールバックを使う必要があります。

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
