---
id: 587d7b89367417b2b2512b49
title: オブジェクトから変数を代入するために分割代入を使う
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

分割代入を使うと、値を取り出すときに新しい変数名を代入できます。値を代入するときにコロンの後に新しい名前を置くことでこれができます。

前の例と同じオブジェクトを使うと：

```js
const user = { name: 'John Doe', age: 34 };
```

代入の際に新しい変数名を付ける方法は次の通りです：

```js
const { name: userName, age: userAge } = user;
```

これは「`user.name` の値を取得して `userName` という新しい変数に代入する」と読み取れます。同様に、`userName` の値は文字列 `John Doe` で、`userAge` の値は数値 `34` になります。

# --instructions--

2つの代入を同等の分割代入に置き換えてください。`highToday` オブジェクトから `highTomorrow` と `today` の値を `tomorrow` と `HIGH_TEMPERATURES` に代入する必要があります。

# --hints--

ES5の代入構文は削除してください。

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

分割代入を使って `highToday` 変数を作成してください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

分割代入を使って `highTomorrow` 変数を作成してください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` は `77` と等しく、`highTomorrow` は `80` と等しくなるべきです。

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
