---
id: 587d7b8d367417b2b2512b59
title: デフォルトエクスポートのインポート
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

前のチャレンジでは`export default`とその使い方について学びました。デフォルトエクスポートをインポートするには、異なる`import`構文を使う必要があります。次の例では、`add`が`math_functions.js`ファイルのデフォルトエクスポートです。インポート方法は以下の通りです。

```js
import add from "./math_functions.js";
```

構文は1つの重要な点で異なります。インポートされる値`add`は波括弧（`{}`）で囲まれていません。ここでの`add`は、`math_functions.js`ファイルのデフォルトエクスポートが何であれ、それを指す変数名です。デフォルトをインポートするときは、ここに任意の名前を使うことができます。

# --instructions--

次のコードで、このファイルと同じディレクトリにある`math_functions.js`ファイルからデフォルトエクスポートをインポートしてください。インポートには`subtract`という名前を付けてください。

# --hints--

`subtract`を`math_functions.js`から正しくインポートする必要があります。

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
