---
id: 587d7b88367417b2b2512b44
title: パラメータ付きのアロー関数を書く
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

通常の関数と同様に、アロー関数にも引数を渡すことができます。

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` は `8` の値を返します。

アロー関数が単一のパラメータを持つ場合、そのパラメータを囲む丸括弧は省略できます。

```js
const doubler = item => item * 2;
```

アロー関数に複数の引数を渡すことも可能です。

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` は `8` の値を返します。

# --instructions--

`myConcat` 関数を書き換えて、`arr2` の内容を `arr1` に追加するようにし、関数がアロー関数の構文を使うようにしてください。

# --hints--

`var` キーワードを置き換える必要があります。

```js
assert.notMatch(code, /var/g);
```

`myConcat` は定数変数にしてください（`const` を使って）。

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat` は2つのパラメータを持つアロー関数であるべきです。

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` は `[1, 2, 3, 4, 5]` を返すべきです。

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

`function` キーワードは使わないでください。

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
