---
id: 587d7b89367417b2b2512b4b
title: 配列から変数に代入するために分割代入を使う
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6では、オブジェクトの分割代入と同じくらい簡単に配列の分割代入ができます。

スプレッド演算子と配列の分割代入の主な違いの一つは、スプレッド演算子は配列のすべての内容をカンマ区切りのリストに展開することです。そのため、どの要素を変数に代入するかを選ぶことはできません。

配列の分割代入では、まさにそれが可能です：

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

コンソールには`a`と`b`の値が`1, 2`として表示されます。

変数`a`には配列の最初の値が代入され、`b`には配列の2番目の値が代入されます。分割代入では、カンマを使って目的のインデックスに到達することで、配列の任意のインデックスの値にもアクセスできます：

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

コンソールには`a`、`b`、`c`の値が`1, 2, 5`として表示されます。

# --instructions--

分割代入を使って`a`と`b`の値を入れ替え、`a`が`b`に格納されている値を受け取り、`b`が`a`に格納されている値を受け取るようにしてください。

# --hints--

入れ替え後、`a`の値は`6`であるべきです。

```js
assert(a === 6);
```

入れ替え後、`b`の値は`8`であるべきです。

```js
assert(b === 8);
```

`a`と`b`を入れ替えるには配列の分割代入を使うべきです。

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
