---
id: 598f48a36c8c40764b4e52b3
title: オブジェクトのミューテーションを防ぐ
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

前のチャレンジで見たように、`const` の宣言だけではデータのミューテーションを本当に防ぐことはできません。データが変更されないようにするために、JavaScript はデータのミューテーションを防ぐ関数 `Object.freeze` を提供しています。

オブジェクトを変更しようとすると、スクリプトが strict モードで実行されている場合はエラーがスローされ、変更は拒否されます。

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

`obj.review` と `obj.newProp` の代入はエラーになります。これはエディターがデフォルトで strict モードで動作しているためで、コンソールには `{ name: "FreeCodeCamp", review: "Awesome" }` の値が表示されます。

# --instructions--

このチャレンジでは、数学定数が変更されないように `Object.freeze` を使います。`MATH_CONSTANTS` オブジェクトを凍結して、`PI` の値を誰も変更できず、プロパティの追加や削除もできないようにしてください。

# --hints--

`const` キーワードは置き換えないでください。

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` は定数変数（`const` を使って）にする必要があります。

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

`MATH_CONSTANTS` の元の宣言は変更しないでください。

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` は `3.14` と等しくなるようにしてください。

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
