---
id: 5cdafbd72913098997531681
title: thenでFulfilledなPromiseを処理する
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

Promiseは、コード内で処理にかかる時間が不明な場合（つまり非同期の何か）、特にサーバーリクエストの際に最も役立ちます。サーバーリクエストを行うと一定の時間がかかり、完了後に通常はサーバーからのレスポンスを使って何かをしたいものです。これは`then`メソッドを使うことで実現できます。

```js
Promise.prototype.then(onFulfilled, onRejected)
```

`then`メソッドは、Promiseの最終的な完了（成功または拒否）に対してコールバック関数をスケジュールします。`onFulfilled`および`onRejected`のハンドラのいずれかが実行され、現在のPromiseの成功または拒否を処理します。Promiseが`resolve`で成功した場合、`onFulfilled`ハンドラが呼び出されます。

```js
myPromise.then(result => {
  
});
```

`result`は`resolve`メソッドに渡された引数から来ています。

# --instructions--

Promiseに`then`メソッドを追加してください。コールバック関数のパラメータとして`result`を使い、`result`をコンソールにログ出力してください。

# --hints--

Promiseに`then`メソッドを呼び出す必要があります。

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

`then`メソッドには、`result`をパラメータとするコールバック関数が必要です。

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

`result`をコンソールにログ出力してください。

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
