---
id: 5cdafbb0291309899753167f
title: JavaScriptのPromiseを作成する
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

JavaScriptのPromiseは、その名前の通り「何かをする約束」を表します。通常は非同期で使います。タスクが完了すると、約束を果たすか失敗します。`Promise`はコンストラクタ関数なので、`new`キーワードを使って作成する必要があります。引数として関数を受け取り、その関数は`resolve`と`reject`という2つのパラメータを持ちます。これらはPromiseの結果を決定するためのメソッドです。構文は次のようになります。

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

`makeServerRequest`という新しいPromiseを作成してください。コンストラクタには`resolve`と`reject`というパラメータを持つ関数を渡してください。

# --hints--

宣言済みの変数`makeServerRequest`にPromiseを代入してください。

```js
assert(makeServerRequest instanceof Promise);
```

Promiseには`resolve`と`reject`をパラメータに持つ関数を渡してください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
