---
id: 5cddbfd622f1a59093ec611d
title: モジュールスクリプトを作成する
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScriptは、もともと主にHTMLで構成されたウェブの中で小さな役割を果たして始まりました。現在では非常に大きくなり、JavaScriptだけでほぼ完全に構築されたウェブサイトもあります。JavaScriptをよりモジュール化し、きれいで保守しやすくするために、ES6ではJavaScriptファイル間でコードを簡単に共有する方法が導入されました。これは、ファイルの一部を他の1つ以上のファイルで使えるようにエクスポートし、必要な部分を必要な場所でインポートすることを含みます。この機能を活用するには、HTMLドキュメント内で`type`が`module`のスクリプトを作成する必要があります。例を示します。

```html
<script type="module" src="filename.js"></script>
```

この`module`タイプを使うスクリプトは、今後のチャレンジで学ぶ`import`と`export`の機能を使えるようになります。

# --instructions--

HTMLドキュメントに`module`タイプのスクリプトを追加し、ソースファイルを`index.js`にしてください。

# --hints--

`script`タグを作成してください。

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

`script`タグには、値が`type`の`module`属性を付けてください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

`script`タグには`src`の`index.js`を付けてください。

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
