---
id: 5cfa550e84205a357704ccb6
title: オブジェクトから値を抽出するために分割代入を使う
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>分割代入</dfn>はES6で導入された特別な構文で、オブジェクトから直接取り出した値をきれいに代入するためのものです。

次のES5のコードを考えてみましょう。

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name`には文字列`John Doe`が代入され、`age`には数値`34`が代入されます。

こちらはES6の分割代入構文を使った同等の代入文です。

```js
const { name, age } = user;
```

同様に、`name`には文字列`John Doe`が代入され、`age`には数値`34`が代入されます。

ここで、`name`と`age`の変数が作成され、`user`オブジェクトのそれぞれの値が代入されます。どれだけすっきりしているかがわかります。

オブジェクトから抽出する値は、必要なだけ多くても少なくてもかまいません。

# --instructions--

2つの代入文を同等の分割代入に置き換えてください。`today`と`tomorrow`の変数に、`today`オブジェクトから`tomorrow`と`HIGH_TEMPERATURES`の値が代入されるようにしてください。

# --hints--

ES5の代入構文は削除してください。

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

分割代入を使って`today`変数を作成してください。

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

分割代入を使って`tomorrow`変数を作成してください。

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today`は`77`と等しく、`tomorrow`は`80`と等しくなるはずです。

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
