---
id: 65688b5a1655a7a6caede847
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Якщо ви спробуєте змінити окремі символи рядка, як це було зроблено в попередньому кроці, ви отримаєте помилку `TypeError`, яка виникає, коли у коді використано об’єкт невідповідного типу.

Як видно з повідомлення помилки, рядки не підтримують призначення елементів, оскільки вони незмінні. However, a variable can be reassigned another string:

```py
message = 'Hello World'
message = 'Hello there!'
```

Delete the `text[0]` line and reassign `text` the string `'Albatross'`.

# --hints--

You should not have `text[0] = 'C'` in your code.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

You should reassign `text` the string `'Albatross'`. Do not modify `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
