---
id: 65687d2f8c7ee27b0446cef3
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Код, який виконується на кожній ітерації (розміщений після `:`), становить тіло циклу. Цей код повинен мати відступ. У Python рекомендується використовувати 4 пробіли на рівень відступу. Цей рівень із відступом є блоком коду.

```py
for i in text:
    print(i)
```

В Python використовуються відступи, щоб вказувати блоки коду. Двокрапка в кінці рядка є сигналом того, що слідуватиме новий блок коду з відступом.

Якщо після останньої двокрапки не знайдено блоку з відступом, виконання коду припиняється та видається `IndentationError`. Цей код не показуватиме вивід, а натомість видасть `IndentationError`:

```py
for i in text:
print(i)
```

Give your `for` loop a body by adding a call to `print(i)`. Remember to indent the loop body.

# --hints--

You should add `print(i)` to your `for` loop body. Pay attention to the indentation.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Код містить недійсний синтаксис та/або недійсні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
