---
id: 63b608ebf40c871d960fc004
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

你還需要兩個這樣的 `fieldset` 代碼塊 - 一個用於 `snacks`，一個用於 `exercise`。

# --hints--

你應該在 `form` 中創建第四個 `fieldset` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[3]);
```

你的第四個 `fieldset` 元素應位於第三個 `fieldset` 元素之後。

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.previousElementSibling?.tagName, "FIELDSET");
```

你的第四個 `fieldset` 元素應該將 `id` 設置爲 `snacks`。

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.id, "snacks");
```

你的第四個 `fieldset` 元素應該包含一個 `legend` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('legend'));
```

你的第四個 `legend` 元素應該包含文本 `Snacks`。

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('legend')?.innerText, "Snacks");
```

你的第四個 `fieldset` 元素應該包含一個 `div` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[3]?.querySelector('div'));
```

Your fourth `div` element should have a `class` set to `input-container`.

```js
assert.equal(document.querySelectorAll('form fieldset')[3]?.querySelector('div')?.className, "input-container");
```

你應該在 `form` 中創建第五個 `fieldset` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[4]);
```

你的第五個 `fieldset` 元素應位於第四個 `fieldset` 元素之後。

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.previousElementSibling?.tagName, "FIELDSET");
```

你的第五個 `fieldset` 元素應將 `id` 設置爲 `exercise`。

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.id, "exercise");
```

你的第五個 `fieldset` 元素應該包含一個 `legend` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('legend'));
```

你的第五個 `legend` 元素應該包含文本 `Exercise`。

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('legend')?.innerText, "Exercise");
```

你的第五個 `fieldset` 元素應該包含一個 `div` 元素。

```js
assert.exists(document.querySelectorAll('form fieldset')[4]?.querySelector('div'));
```

你的第五個 `div` 元素應將 `class` 設置爲 `input-container`。

```js
assert.equal(document.querySelectorAll('form fieldset')[4]?.querySelector('div')?.className, "input-container");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--

--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
