---
id: 63b60af1a0b9f7238a9dd294
title: 步驟 10
challengeType: 0
dashedName: step-10
---

# --description--

你的選擇菜單需要爲你在前面的步驟中創建的每種食物和運動 `fieldset` 元素提供選項。 使用 `option` 元素爲每個 `fieldset` 創建一個新選項。 每個選項的 `value` 屬性應爲 `fieldset` 的 `id`，每個選項的文本應爲 `legend` 的文本。

將 `Breakfast` 選項設置爲 `selected` 選項。

# --hints--

你應該在 `select` 元素內創建五個 `option` 元素。

```js
assert.equal(document.querySelectorAll('.controls select option')?.length, 5);
```

你的第一個 `option` 應該包含文本 `Breakfast`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.textContent?.trim(), 'Breakfast');
```

你的第一個 `option` 應該將 `value` 屬性設置爲 `breakfast`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[0]?.value, 'breakfast');
```

你的第二個 `option` 應該包含文本 `Lunch`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.textContent?.trim(), 'Lunch');
```

你的第二個 `option` 應該將 `value` 屬性設置爲 `lunch`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[1]?.value, 'lunch');
```

你的第三個 `option` 應該包含文本 `Dinner`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.textContent?.trim(), 'Dinner');
```

你的第三個 `option` 應將 `value` 屬性設置爲 `dinner`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[2]?.value, 'dinner');
```

你的第四個 `option` 應該包含文本 `Snacks`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.textContent?.trim(), 'Snacks');
```

你的第四個 `option` 應該將 `value` 屬性設置爲 `snacks`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[3]?.value, 'snacks');
```

你的第五個 `option` 應該包含文本 `Exercise`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.textContent?.trim(), 'Exercise');
```

你的第五個 `option` 應該將 `value` 屬性設置爲 `exercise`。

```js
assert.equal(document.querySelectorAll('.controls select option')?.[4]?.value, 'exercise');
```

你的第一個 `option` 應該被默認選擇。

```js
const isFirstOptionSelected = document.querySelectorAll('.controls select option')?.[0]?.getAttributeNames()?.includes('selected');
const selectedOptions = document.querySelectorAll('.controls select option[selected]');
assert.isTrue(isFirstOptionSelected && selectedOptions.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>
              <label for="entry-dropdown">Add food or exercise:</label>
              <select id="entry-dropdown" name="options">

              </select>
              <button type="button" id="add-entry">Add Entry</button>
            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
