---
id: 65687b48f2201ab32e06c37c
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

目前，你的脚本会抛出 `TypeError`，因为你正在尝试将字符串添加到整数。 你可以通过使用内置的 `int` 函数将 `digit` 变量转换为整数，然后将其添加到 `sum_of_odd_digits` 来解决此问题：

```python
my_string = '123'
my_int = int(my_string)
```

将 `digit` 变量转换为整数，然后将其添加到 `sum_of_odd_digits`。 然后，将 `print` 调用移至 `verify_card_number` 函数的末尾，以打印 `sum_of_odd_digits` 的值。

# --hints--

`for` 循环中应该有 `sum_of_odd_digits += int(digit)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

`verify_card_number` 函数末尾应该有 `print(sum_of_odd_digits)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

你不应该再使用 `print(digit)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
