---
id: 656879c6f35a85a59c06b3a7
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

你之前已使用索引运算符 `[]` 访问过字符串的元素（字符）。 你还可以使用索引运算符通过 `string[start:stop:step]` 访问字符串中的一系列字符：

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

其中 `start` 是起始索引（包含），`stop` 是结束索引（不包含），`step` 是要跳过的字符数。 如果未指定，`step` 默认为 1。

创建一个名为 `card_number_reversed` 的变量，并将 `card_number` 的前 4 个字符的值赋给它。

# --hints--

`verify_card_number` 函数中应该有 `card_number_reversed = card_number[0:4]`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
