---
id: 65687ad86596e0af38640a84
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

使用 `for` 循环迭代 `odd_digits` 列表中的每个数字。 将上一步中的 `print` 调用移至 `for` 循环中，并将其更改为打印每个数字。

# --hints--

你应该对 `odd_digits` 使用 `for` 循环。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /for\s+\w+\s+in\s+odd_digits/);
    }
})
```

`for` 循环中应该有 `--fcc-expected--`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;
    
        // Get variable name used in for loop
        const for_loop_variable = function_body.match(/for\s+(\w+)\s+in\s+odd_digits/)?.[1];
        assert.exists(for_loop_variable);

        assert.equal(function_body.match(new RegExp(`print\\(\\s*${for_loop_variable}\\s*\\)`))?.[0], `print(${for_loop_variable})`);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    print(odd_digits)
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
