---
id: 656878585631369a6b2d2191
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

定义翻译本身并不会翻译字符串。 必须对要翻译的字符串调用 `translate` 方法，并将翻译表作为参数：

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

创建一个名为 `translated_card_number` 的变量，并将以 `card_translation` 作为参数对 `card_number` 调用 `translate` 方法的结果分配给它。

# --hints--

你应该在 `main` 中创建一个 `translated_card_number` 变量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

你应该为 `translated_card_number` 分配一个值 `card_number.translate(card_translation)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
