---
id: 656879613a96aba21cbe80b9
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

在 `main` 函数中，调用 `verify_card_number` 函数并传入 `translated_card_number` 变量作为参数。

# --hints--

你应该在 `main` 函数中包含 `verify_card_number(translated_card_number)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +verify_card_number\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

--fcc-editable-region--
main()
```
