---
id: 65687a8253766cac4c99c57f
title: Step 17
challengeType: 20
dashedName: step-17
---

# --description--

在 `verify_card_number` 函数中，创建一个变量 `odd_digits`，其中包含 `card_number_reversed` 字符串的所有其他数字。

# --hints--

`verify_card_number` 函数中应该有 `odd_digits = card_number_reversed[::2]`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /odd_digits\s*=\s*card_number_reversed\s*\[\s*:\s*:\s*2\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
