---
id: 65687b182c2a8fb1470e0b97
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

在 `for` 循环中，使用 `+=` 运算符将 `digit` 添加到 `sum_of_odd_digits` 变量。

执行此操作时，你的脚本会抛出 `TypeError`，因为你正在尝试将字符串添加到整数，但不用担心，你将在下一步中了解有关如何使其工作的更多信息。

# --hints--

`for` 循环中应该有 `sum_of_odd_digits += digit`。

```js
const transformedCode = code.replace(/\r/g, "");
const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
const { function_body } = verify_card_number;

assert.match(function_body, /sum_of_odd_digits\s*\+=\s*digit/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        print(digit)
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
