---
id: 656877c202f44995ed952608
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

将 `pass` 语句替换为名为 `card_number` 的变量和值 `'4111-1111-4555-1142'`。

# --hints--

`main` 函数中应该有 `card_number = '4111-1111-4555-1142'`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

你的代码中不应包含 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
