---
id: 656879a66338b2a461d5d307
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

Luhn算法如下：

1. 从右到左，每隔一个数字的值加倍；如果乘积大于 9，则将乘积的数字相加。
2. 取所有数字的总和。
3. 如果所有数字之和是10的倍数，则该数字有效；否则无效。

假设一个帐号“7992739871”的示例将添加一个校验位，使其具有 7992739871x 的形式：

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

将 `pass` 语句替换为变量 `sum_of_odd_digits` 和值 `0`。

# --hints--

`verify_card_number` 函数中应该有 `sum_of_odd_digits = 0`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

你不应在 `verify_card_number` 函数中使用 `pass`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
