---
id: 65687b8eb0bbf7b5d41b610b
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

创建一个变量 `even_digits` 并为其分配反转卡号的偶数位。

# --hints--

`verify_card_number` 函数中应该有 `even_digits = card_number_reversed[1::2]`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /even_digits\s*=\s*card_number_reversed\s*\[\s*1\s*:\s*:\s*2\s*\]/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)
    print(sum_of_odd_digits)

--fcc-editable-region--
    sum_of_even_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
