---
id: 656877f71bba2b97acedf9af
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

Python 带有内置的类，可以帮助我们进行字符串操作。 其中之一是 `str` 类。 它有一个名为 `maketrans` 的方法可以帮助我们创建翻译表。 该表可用于替换字符串中的字符：

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

当在字符串上调用时，上面的代码将用 `c` 替换所有 `t` 字符，用 `b` 替换所有 `l` 字符。

创建一个名为 `card_translation` 的变量，并赋值一个翻译表，用于将所有 `-` 和 ` ` 字符替换为空字符串。

# --hints--

你应该在 `main` 中创建一个 `card_translation` 变量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

你应该为 `card_translation` 分配一个值 `str.maketrans({'-': '', ' ': ''})`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
