---
id: 65687db6a9fcf3c548a6c1b7
title: 第 28 步
challengeType: 20
dashedName: step-28
---

# --description--

该算法的一部分是从右侧开始，每隔两个数字加倍。 如果数字加倍的结果大于或等于`10`，则将这两个数字相加。 例如，如果数字是 `6`，则将其加倍得到 `12`。 将 `1` 和 `2` 加在一起得到 `3`。 你可以通过使用*整数除法*来获取第一个数字，并使用模运算符 (`%`) 来获取第二个数字：

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

整数除法得到除法的商，向下舍入到最接近的整数。

在 `if` 语句中，将 `number // 10`（整数除法）的结果加上 `number` 的模赋予 `number` > 和 `10`。

# --hints--

`if` 语句中应包含 `number = number // 10 + number % 10`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
