---
id: 65687da2e60409c45595bbe1
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

为了防止一位数字的乘积大于 `9`，请在偶数数字循环中添加一个 `if` 语句来检查 `number` 是否大于小于或等于`10`。 如果是，则打印`number`。

# --hints--

`for` 循环中应该有 `if number >= 10:`。

```js
({
    test: () => 
    {
      assert(runPython(`_Node(_code).find_function("verify_card_number").find_for_loops()[1].find_ifs()[0].find_conditions()[0].is_equivalent("number >= 10")`));
    }
})
```

你应该在 `if` 语句中包含 `print(number)`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const [_,lower_function_body] = function_body.split("sum_of_even_digits");

        assert.match(lower_function_body, /print\s*\(\s*number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
