---
id: 6569b831a07d7154c793301b
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

定义一个带有参数 `card_number` 的函数 `verify_card_number`，并使用 `pass` 关键字使该函数不执行任何操作。

# --hints--

你应该使用一个名为 `verify_card_number` 的函数，其参数名为 `card_number`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

你应该在 `verify_card_number` 函数主体中使用 `pass` 关键字。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
