---
id: 65c4f09e074dd8d37830ea00
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Crea un'altra classe vuota chiamata `BinarySearchTree` che rappresenta un albero binario di ricerca.

# --hints--

Dovresti avere una classe chiamata `BinarySearchTree`. Ricorda di usare la parola chiave `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
