---
id: 65c4f2d9fd872fd99ac659dd
title: Passaggio 13
challengeType: 20
dashedName: step-13
---

# --description--

Aggiungi un'istruzione condizionale `elif` che verifichi se `key` > `node.key`.

# --hints--

Dovresti aggiungere un'istruzione condizionale `elif` per controllare se `key > node.key`. Ricorda di usare la parola chiave `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[1].is_equivalent("key > node.key")`))})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)

--fcc-editable-region--    

```
