---
id: 65c4f0b82db314d3fcc0b8e1
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

All'interno della classe `BinarySearchTree`, sostituisci `pass` con un metodo `__init__` e aggiungi un parametro `self` a questo metodo.

# --hints--

Dovresti rimuovere la parola chiave `pass` dalla classe `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Dovresti definire un metodo `__init__` e aggiungere un parametro `self` a questo metodo. Ricorda di usare la parola chiave `pass` all'interno del corpo del metodo.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
