---
id: 65c4f2b7178afed8e88f782f
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Se `key` < `node.key` valuta a `True`, allora il nuovo nodo dovrebbe essere posizionato nel sottoalbero sinistro.

Elimina `pass` e richiama ricorsivamente il metodo `_insert` con il figlio sinistro come primo argomento e `key` come secondo argomento. Assegna il risultato all'attributo `left` del nodo corrente.

# --hints--

Dovresti rimuovere la parola chiave `pass` dal blocco `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Dovresti chiamare il metodo `self._insert` passando `node.left` e `key` come argomenti.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Dovresti assegnare il risultato della tua chiamata `self._insert()` all'attributo `left` del nodo corrente.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
