---
id: 65c4f22498d22ed775ef8efb
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Ora devi verificare se il parametro `node` è `None`. Se lo è, significa che il metodo ha raggiunto un nodo foglia o uno spazio vuoto nell'albero dove dovrebbe essere inserito il nuovo nodo.

All'interno del corpo del metodo `_insert`, sostituisci `pass` con un'istruzione `if` che verifica se `node is None`.

All'interno del nuovo blocco `if`, restituisci `TreeNode(key)` per creare una nuova istanza `TreeNode` con la chiave fornita. Questo diventerà il nuovo nodo foglia, inserendo di fatto la chiave nell'albero.

# --hints--

Non dovresti avere `pass` nel tuo metodo `_insert`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Dovresti scrivere un'istruzione `if` per verificare se `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Dovresti restituire `TreeNode(key)` dal blocco `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
