---
id: 65c4f02dbd3557d23f12db62
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

All'interno del metodo `__init__`, elimina `pass` e assegna il valore del parametro `key` all'attributo `key` del nodo usando `self.key`.

Questo significa che l'attributo `key` dell'istanza `TreeNode` sarà impostato al valore passato durante la creazione dell'oggetto.

# --hints--

Dovresti eliminare l'istruzione `pass` dal metodo `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Dovresti assegnare il valore del parametro `key` all'attributo `key` del nodo usando `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
