---
id: 65c4f2851645e8d84d14f5d1
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ora devi attraversare ricorsivamente l'albero e inserire i valori seguendo il principio per gli alberi binari:

- I valori minori della chiave vanno nel sottoalbero sinistro
- I valori maggiori della chiave vanno nel sottoalbero destro

Dopo la tua istruzione condizionale esistente, scrivi un'altra istruzione `if` per verificare se `key` è minore di `node.key`.

# --hints--

Dovresti scrivere un'altra istruzione if per verificare se `key` è minore di `node.key`. Ricorda di usare la parola chiave `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```
