---
id: 65c4eb814cc977c95cd7df0e
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto creerai un albero binario di ricerca (BST). Un BST è una struttura dati in cui ogni nodo ha al massimo due figli, con il figlio sinistro che contiene valori minori del nodo genitore e il figlio destro che contiene valori maggiori del nodo genitore, permettendo operazioni efficienti di ricerca e ordinamento.

Ecco come appare un albero binario di ricerca:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Inizia definendo una classe `TreeNode` vuota. La classe `TreeNode` rappresenta un nodo in un albero binario di ricerca. Usa la parola chiave `pass` per completare il corpo della classe ed evitare un errore.

# --hints--

Dovresti usare la parola chiave `class` per dichiarare una classe vuota chiamata `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
