---
id: 65c4f013851cefd1a4fe4c96
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Il metodo `__init__` prende due parametri: `self` (che rappresenta l’istanza della classe che si sta creando) e `key` (il valore da archiviare nel nodo). Aggiungi questi due parametri al metodo `__init__()`.

# --hints--

Dovresti aggiungere il parametro `self` al metodo.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Dovresti aggiungere il parametro `key` al metodo.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
