---
id: 65c4f07da204f4d2f325fbd3
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

All'interno del metodo `__init__`, inizializza gli attributi `left` e `right` del nodo a `None`. Questo perché quando un nodo viene creato per la prima volta, non ha figli sinistro o destro. Ricorda di usare la parola chiave `self`. 

# --hints--

Dovresti assegnare `None` all'attributo `left` del nodo usando `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Dovresti assegnare `None` all'attributo `right` del nodo usando `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
