---
id: 65c4f33bc3c3d8db7f732787
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Ora, per eseguire l'inserimento vero e proprio, definisci un metodo `insert` vuoto all'interno della classe `BinarySearchTree` e dagli un parametro `self`.

# --hints--

Dovresti definire un metodo `insert` con un parametro `self` all'interno della classe `BinarySearchTree`. Ricorda la parola chiave `pass`.

```js
({ test: () => assert.match(code, /^\s{4}def\s+insert\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node
--fcc-editable-region--    

--fcc-editable-region--

```
