---
id: 65c4f195de7e2ad5932be717
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Ora devi definire un meccanismo per inserire nodi nell'albero. Per questo, devi definire un metodo `_insert`, che è una funzione di supporto e verrà usata successivamente dal metodo `insert` vero e proprio.

Questo metodo è ricorsivo, cioè si chiama da solo per attraversare l'albero finché non trova la posizione giusta per il nuovo nodo.

Definisci un metodo `_insert` con i parametri `self`, `node` e `key`.

# --hints--

Dovresti definire un metodo `_insert` all'interno della classe `BinarySearchTree`. Ricorda di usare `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Il tuo metodo `_insert` dovrebbe prendere tre parametri: `self`, `node` e `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
