---
id: 65c4f3258d2e4cdacc919dfd
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Alla fine del tuo metodo `_insert`, dopo che il processo di inserimento è completo, restituisci il nodo corrente per aggiornare la struttura dell'albero ai livelli superiori dello stack di chiamate ricorsive.


# --hints--

Dovresti restituire il nodo corrente all'esterno dei blocchi condizionali.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```
