---
id: 65c4ef8463f869d0eaf87c67
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

All'interno della classe `TreeNode`, sostituisci `pass` con un metodo `__init__` in modo da poter inizializzare gli attributi dell'oggetto. Per ora non aggiungere parametri.

# --hints--

Dovresti rimuovere la parola chiave `pass` dalla classe `TreeNode` e spostarla all'interno del metodo `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Dovresti definire un metodo `__init__` all'interno della tua classe `TreeNode`. Ricorda di usare la parola chiave `pass` all'interno.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
