---
id: 65c4f300da28d8da361bfa93
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

All'interno della tua clausola `elif`, sostituisci `pass` con una chiamata al metodo `_insert` con il figlio destro del nodo corrente come primo argomento e `key` come secondo argomento. Assegna il risultato all'attributo `right` del nodo corrente.

# --hints--

Dovresti rimuovere la parola chiave `pass` dal blocco `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Dovresti chiamare il metodo `self._insert()` passando `node.right` e `key` come argomenti.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Dovresti assegnare il risultato della tua chiamata `self._insert()` all'attributo corretto (`node.right`) del nodo corrente.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
