---
id: 65c4f37ba6cbcfdc77e24165
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

La funzione `insert` sarà chiamata dall'utente. Oltre al parametro `self`, avrà bisogno anche di un parametro `key`. Questo parametro sarà il valore chiave da inserire nell'albero binario di ricerca.

Aggiungi `key` come secondo parametro alla definizione della funzione.

# --hints--

La funzione `insert` dovrebbe contenere il parametro `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

La funzione `insert` dovrebbe contenere il parametro `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
