---
id: 65c4f0d32ca17ad4b1636b0e
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

All'interno del metodo `__init__`, elimina `pass` e inizializza `root` al valore `None`.

L'attributo `root` rappresenta il nodo radice dell'albero binario di ricerca. Poiché questo è il costruttore, quando viene creato un nuovo oggetto `BinarySearchTree`, inizia con un albero vuoto, quindi l'attributo `root` viene impostato su `None`.

# --hints--

Dovresti rimuovere l'istruzione `pass` dal metodo `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Dovresti inizializzare l'attributo `root` a `None` usando `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
