---
id: 65c6409418806da8d0636ffc
title: Passaggio 26
challengeType: 20
dashedName: step-26
---

# --description--

Le funzionalità `insert` e `search` sono complete, è ora di testarle.

Nota che, a questo punto, i nodi non sono ordinati ma solo inseriti. Lavorerai sull'ordinamento usando la visita inorder più avanti.

Puoi creare un'istanza di una classe in Python così:

```python
object_name = ClassName()
```

Fuori dalle definizioni delle classi, crea un'istanza della classe `BinarySearchTree` e assegnala alla variabile `bst`.


# --hints--

Dovresti creare un'istanza della classe `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Dovresti assegnare la nuova istanza di `BinarySearchTree` alla variabile `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
