---
id: 65c63df529bd15a24c187c62
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Scrivi un'altra istruzione `if` che verifica se la chiave obiettivo è minore della chiave del nodo corrente.

All'interno del blocco `if`, restituisci il risultato della chiamata al metodo `_search` con il figlio sinistro del nodo corrente e `key` come argomenti.

# --hints--

Dovresti scrivere un'altra istruzione `if` per verificare se `key` < `node.key`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

Dovresti restituire `self._search(node.left, key)` dal tuo nuovo blocco `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_search").find_ifs()[1].find_bodies()[0].is_equivalent("return self._search(node.left, key)")`)) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node

--fcc-editable-region--
```
