---
id: 65c63e6962159fa372ecdbec
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

Se la seconda istruzione `if` non è `True`, significa che la chiave obiettivo è maggiore o uguale alla chiave del nodo corrente.

In un albero binario di ricerca, se la chiave obiettivo è maggiore della chiave del nodo corrente, la ricerca continua nel sottoalbero destro.

Dopo il blocco `if`, restituisci il risultato della chiamata al metodo `_search` con il figlio destro del nodo corrente e la chiave come argomenti.

# --hints--

Dovresti chiamare ricorsivamente il metodo `_search` usando `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Dovresti passare due argomenti alla chiamata del metodo `_search`: `node.right` e `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Dovresti restituire il risultato della chiamata al metodo `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
