---
id: 65c63da2ed6769a10e141341
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Ora definirai un caso base per la ricerca ricorsiva. Rimuovi l'attuale `pass` e scrivi un'istruzione `if` che verifichi due condizioni:

- Se `node` è `None`: questo indica che la ricerca ha raggiunto la fine di un ramo senza trovare la chiave.
- Se `node.key == key`: significa che la chiave è stata trovata nel nodo corrente.

Combina le due condizioni con l'operatore `or` e restituisci il nodo corrente all'interno del blocco `if`.

# --hints--

Dovresti rimuovere la parola chiave `pass` dal metodo `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Dovresti scrivere un'istruzione `if` che verifichi se `node` è `None` oppure se `node.key` è uguale a `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Dovresti restituire `node` dal blocco `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
