---
id: 65c63fd9b387f0a6c5bf6a72
title: Passaggio 24
challengeType: 20
dashedName: step-24
---

# --description--

All'interno del metodo `search`, elimina `pass` e chiama il metodo di supporto `_search` con i seguenti argomenti.

- `self.root`: questa è la radice dell'albero binario di ricerca. La ricerca inizia dalla radice.
- `key`: questo è il valore che l'utente vuole trovare nell'albero binario di ricerca.

Internamente, `search` delega la logica effettiva della ricerca al metodo di supporto `_search` che esegue la ricerca ricorsiva vera e propria all'interno dell'albero binario di ricerca.

# --hints--

Dovresti rimuovere la parola chiave `pass` dal metodo `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Dovresti chiamare il metodo `_search` all'interno del metodo `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Dovresti chiamare il metodo `_search` passando `self.root` e `key` come argomenti.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
