---
id: 65c63a4c4da62e9ae18e321a
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

È ora di lavorare sulla funzionalità di ricerca. Proprio come hai creato un metodo helper `_insert` per il metodo `insert`, devi creare un metodo helper `_search` per il metodo `search`.

Definisci il metodo `_search` con tre parametri, cioè `self`, `node` e `key`.

# --hints--

Dovresti dichiarare un metodo `_search`. Nota l'underscore iniziale e ricorda di usare `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Il tuo metodo `_search` dovrebbe prendere tre parametri: `self`, `node` e `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
