---
id: 65cb45d130c97cb459439fac
title: Passaggio 42
challengeType: 20
dashedName: step-42
---

# --description--

Per trovare il valore più piccolo nel sottoalbero destro, devi iterare sui figli sinistri del nodo dato finché non raggiungi il nodo più a sinistra (il più piccolo) nel sottoalbero.

Per farlo, scrivi un loop `while` che si esegue quando `node.left is not None` e sposta `pass` all'interno del blocco `while`. Questa condizione verifica se esiste un figlio sinistro. Finché c'è un figlio sinistro, il loop continua e c'è un valore più piccolo da trovare.

# --hints--

dovresti rimuovere la parola chiave `pass` dal metodo `_min_value`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_min_value").has_pass()`
      )
    );
  },
});
```

Dovresti usare la condizione `node.left is not None` nel loop `while`.

```js
({ test: () => assert.match(code, /while\s+node\.left\s+is\s+not\s+None\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left  

--fcc-editable-region--
    def _min_value(self, node):
        pass

--fcc-editable-region--   

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
