---
id: 69008059d43d6751385883c1
title: Depura errores del operador de incremento y decremento en una aplicación con fallos
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Tu amigo te ha pedido que le ayudes a depurar su código. Está practicando los operadores de incremento y decremento, pero está obteniendo resultados inesperados en la consola.

Tu tarea es ayudarles a identificar los errores y actualizar el código para que las salidas en la consola coincidan con los resultados esperados.

Puedes referirte a las lecciones anteriores sobre operadores de incremento y decremento, así como a la notación prefija y posfija.

**Historias de usuario:**

1. Debes tener una variable llamada `updatedCounter` que produzca correctamente el valor `11` usando el operador de incremento. Actualiza la expresión para usar la notación prefija o sufija correcta en lugar de codificar el valor directamente.
2. Debes tener una variable llamada `finalScore` que produzca correctamente el valor `8` usando el operador de incremento. Actualiza la expresión para usar la notación prefijo o sufijo correcta en lugar de codificar el valor directamente.
3. Debes tener una variable llamada `updatedCoins` que produzca correctamente el valor `2` usando el operador de decremento. Actualiza la expresión para usar la notación prefija o posfija correcta en lugar de codificar el valor directamente.
4. Debes tener una variable llamada `newHealth` que produzca correctamente el valor `7` usando el operador de decremento. Actualiza la expresión para usar la notación prefijo o sufijo correcta en lugar de codificar el valor directamente.

# --hints--

Debes tener una variable llamada `updatedCounter` que use la notación de incremento prefijo (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` debería ser igual a `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Debes tener una variable llamada `finalScore` que use la notación de incremento posfijo (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` debería ser igual a `8`.

```js
assert.strictEqual(finalScore, 8);
```

Debes tener una variable llamada `updatedCoins` que use la notación de decremento prefijo (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` debería ser igual a `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Debes tener una variable llamada `newHealth` que use la notación de decremento posfijo (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` debería ser igual a `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
