---
id: 587d8249367417b2b2512c40
title: Налаштування Helmet за допомогою батьківського підпрограмного забезпечення helmet()
challengeType: 2
forumTopicId: 301575
dashedName: configure-helmet-using-the-parent-helmet-middleware
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

`app.use(helmet())` автоматично включатиме все підпрограмне забезпечення, представлене вище, за винятком `noCache()` та `contentSecurityPolicy()`, але їх можна увімкнути за потреби. Ви також можете індивідуально вимкнути або налаштувати будь-яке інше підпрограмне забезпечення, використовуючи об’єкт конфігурації.

**Приклад:**

```js
app.use(helmet({
  frameguard: {         // configure
    action: 'deny'
  },
  contentSecurityPolicy: {    // enable and configure
    directives: {
      defaultSrc: ["'self'"],
      styleSrc: ['style.com'],
    }
  },
  dnsPrefetchControl: false     // disable
}))
```

Ми представляли кожне підпрограмне забезпечення окремо для навчальних цілей і для зручності тестування. Використання батьківського `helmet()` підпрограмного забезпечення легко реалізувати в реальному проєкті.

# --hints--

немає тестів — це описове завдання

```js
assert(true);
```
