---
id: 587d8249367417b2b2512c3e
title: Вимкнути кешування на стороні клієнта за допомогою helmet.noCache()
challengeType: 2
forumTopicId: 301576
dashedName: disable-client-side-caching-with-helmet-nocache
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Якщо ви випускаєте оновлення для свого вебсайту і хочете, щоб користувачі завжди скачували новішу версію, ви можете спробувати вимкнути кешування в браузері клієнта. Це також може бути корисним під час розробки. Кешування має переваги для продуктивності, які ви втратите, тому використовуйте цей варіант лише тоді, коли це справді необхідно.

# --instructions--

Використайте метод `helmet.noCache()` на своєму сервері.

# --hints--

Підпрограмне забезпечення helmet.noCache() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nocache');
assert.equal(
  data.headers['cache-control'],
  'no-store, no-cache, must-revalidate, proxy-revalidate'
);
```
