---
id: 587d8248367417b2b2512c3d
title: Вимкнути попереднє завантаження DNS за допомогою helmet.dnsPrefetchControl()
challengeType: 2
forumTopicId: 301577
dashedName: disable-dns-prefetching-with-helmet-dnsprefetchcontrol
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Щоб покращити продуктивність, більшість браузерів попередньо завантажують DNS-записи для посилань на сторінці. Таким чином, IP-адреса призначення вже відома, коли користувач натискає на посилання. Це може призвести до надмірного навантаження на DNS-сервіс (якщо у вас великий вебсайт, який відвідують мільйони людей…), проблем із конфіденційністю (один підслуховувач може дізнатися, що ви перебуваєте на певній сторінці) або спотворення статистики сторінок (деякі посилання можуть виглядати як відвідані, навіть якщо це не так). Якщо у вас високі вимоги до безпеки, ви можете вимкнути попереднє завантаження DNS, але це призведе до зниження продуктивності.

# --instructions--

Використайте метод `helmet.dnsPrefetchControl()` на своєму сервері.

# --hints--

Підпрограмне забезпечення helmet.dnsPrefetchControl() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'dnsPrefetchControl');
assert.equal(data.headers['x-dns-prefetch-control'], 'off');
```
