---
id: 587d8248367417b2b2512c3b
title: Запобігання відкриттю ненадійного HTML в IE за допомогою helmet.ieNoOpen()
challengeType: 2
forumTopicId: 301584
dashedName: prevent-ie-from-opening-untrusted-html-with-helmet-ienoopen
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Деякі веб-застосунки пропонують скачати ненадійний HTML. Деякі версії Internet Explorer за замовчуванням відкривають ці HTML-файли в контексті вашого сайту. Це означає, що ненадійна HTML-сторінка може почати виконувати шкідливі дії в контексті ваших сторінок. Це підпрограмне забезпечення встановлює заголовок X-Download-Options зі значенням noopen. Це запобігатиме виконанню завантажень користувачами IE у контексті довіреного сайту.

# --instructions--

Використайте метод `helmet.ieNoOpen()` на своєму сервері.

# --hints--

Підпрограмне забезпечення helmet.ieNoOpen() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'ienoopen');
assert.equal(data.headers['x-download-options'], 'noopen');
```
