---
id: 587d8248367417b2b2512c3c
title: Запитайте браузери використовувати ваш сайт лише через HTTPS за допомогою helmet.hsts()
challengeType: 2
forumTopicId: 301573
dashedName: ask-browsers-to-access-your-site-via-https-only-with-helmet-hsts
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

HTTP Strict Transport Security (HSTS) — це політика веб-безпеки, яка допомагає захистити сайти від атак пониження протоколу та викрадення cookie. Якщо ваш сайт доступний через HTTPS, ви можете попросити браузери користувачів уникати використання небезпечного HTTP. Встановлюючи заголовок Strict-Transport-Security, ви повідомляєте браузерам використовувати HTTPS для майбутніх запитів протягом певного часу. Це працюватиме для запитів, що надходять після ініціального запиту.

# --instructions--

Налаштуйте `helmet.hsts()`, щоб використовувати HTTPS протягом наступних 90 днів. Передайте об’єкт конфігурації `{maxAge: timeInSeconds, force: true}`. Ви можете створити змінну `ninetyDaysInSeconds = 90*24*60*60;`, щоб використати її для `timeInSeconds`.

Примітка: налаштування HTTPS на кастомному сайті вимагає придбання домену та SSL/TLS сертифікату.

# --hints--

Підпрограмне забезпечення helmet.hsts() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hsts');
assert.property(data.headers, 'strict-transport-security');
```

maxAge має дорівнювати 7776000 с (90 днів)

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.match(
  data.headers['strict-transport-security'],
  /^max-age=7776000;?/
);
```
