---
id: 587d8247367417b2b2512c38
title: Зменшення ризику clickjacking за допомогою helmet.frameguard()
challengeType: 2
forumTopicId: 301582
dashedName: mitigate-the-risk-of-clickjacking-with-helmet-frameguard
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Ваша сторінка може бути розміщена у `<frame>` або `<iframe>` без вашої згоди. Це може призвести до атак clickjacking, серед іншого. Clickjacking — це техніка обману користувача, щоб він взаємодіяв зі сторінкою, відмінною від тієї, яку він вважає. Це досягається шляхом виконання вашої сторінки в шкідливому контексті за допомогою iframing. У такому контексті хакер може накласти прихований шар поверх вашої сторінки. Приховані кнопки можуть запускати шкідливі скрипти. Це підпрограмне забезпечення встановлює заголовок X-Frame-Options. Воно обмежує, хто може розміщувати ваш сайт у фреймі. Існує три режими: DENY, SAMEORIGIN і ALLOW-FROM.

Нам не потрібно, щоб наш застосунок був у фреймі.

# --instructions--

Використайте `helmet.frameguard()`, передаючи об’єкт конфігурації `{action: 'deny'}`.

# --hints--

Підпрограмне забезпечення helmet.frameguard() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(
  data.appStack,
  'frameguard',
  'helmet.frameguard() middleware is not mounted correctly'
);
```

Для helmet.frameguard() параметр 'action' має бути встановлений у 'DENY'

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.property(data.headers, 'x-frame-options');
assert.equal(data.headers['x-frame-options'], 'DENY');
```
