---
id: 587d8248367417b2b2512c3a
title: Уникнення вгадування MIME-типу відповіді за допомогою helmet.noSniff()
challengeType: 2
forumTopicId: 301574
dashedName: avoid-inferring-the-response-mime-type-with-helmet-nosniff
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Браузери можуть використовувати вгадування вмісту або MIME-типу, щоб замінити заголовок `Content-Type` у відповіді і обробити дані, використовуючи неявний тип вмісту. Хоча це може бути зручним у деяких випадках, це також може призвести до небезпечних атак. Це підпрограмне забезпечення встановлює заголовок `X-Content-Type-Options` у значення `nosniff`, вказуючи браузеру не ігнорувати наданий `Content-Type`.

# --instructions--

Використайте метод `helmet.noSniff()` на своєму сервері.

# --hints--

Підпрограмне забезпечення helmet.noSniff() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'nosniff');
assert.equal(data.headers['x-content-type-options'], 'nosniff');
```
