---
id: 587d8247367417b2b2512c36
title: Встановлення та підключення Helmet
challengeType: 2
forumTopicId: 301581
dashedName: install-and-require-helmet
---

# --description--

Робота над цими завданнями передбачає написання вашого коду одним із таких способів:

- Клонуйте <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій GitHub</a> і виконайте завдання локально.
- Використайте будь-який конструктор сайтів на свій вибір для виконання проєкту. Обов’язково включіть усі файли з нашого репозиторію GitHub.

Helmet допомагає захистити ваші застосунки Express, встановлюючи різні HTTP-заголовки.

# --instructions--

Увесь ваш код для цих уроків розміщуйте у файлі `myApp.js` між рядками коду, які ми для вас підготували. Не змінюйте і не видаляйте код, який ми додали.

Версія Helmet `3.21.3` вже встановлена, тому підключіть її як `helmet` у `myApp.js`.

# --hints--

Версія `helmet` `3.21.3` має бути у `package.json`

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
const packJson = JSON.parse(data);
const helmet = packJson.dependencies.helmet;
assert(helmet === '3.21.3' || helmet === '^3.21.3');
```
