---
id: 587d8247367417b2b2512c37
title: Приховати потенційно небезпечну інформацію за допомогою helmet.hidePoweredBy()
challengeType: 2
forumTopicId: 301580
dashedName: hide-potentially-dangerous-information-using-helmet-hidepoweredby
---

# --description--

Нагадаємо, цей проєкт створюється на основі шаблонного коду, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-infosec/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Хакери можуть скористатися відомими вразливостями Express/Node, якщо побачать, що ваш сайт працює на Express. `X-Powered-By: Express` надсилається у кожному запиті, що надходить від Express, за замовчуванням. Використайте підпрограмне забезпечення `helmet.hidePoweredBy()`, щоб видалити заголовок X-Powered-By.

# --hints--

Підпрограмне забезпечення helmet.hidePoweredBy() має бути правильно підключене

```js
const response = await fetch(code + '/_api/app-info');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.json();
assert.include(data.appStack, 'hidePoweredBy');
assert.notEqual(data.headers['x-powered-by'], 'Express');
```
