---
id: 58a25bcff9fc0f352b528e7e
title: Хешування та порівняння паролів синхронно
challengeType: 2
forumTopicId: 301579
dashedName: hash-and-compare-passwords-synchronously
---

# --description--

Нагадаємо, цей проєкт базується на стартовому проєкті, скопійованому з <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Синхронне хешування так само легко виконати, але воно може спричиняти затримки, якщо використовувати його на сервері з високою вартістю або при дуже частому хешуванні. Хешування цим методом таке ж просте, як виклик

```js
var hash = bcrypt.hashSync(myPlaintextPassword, saltRounds);
```

Додайте цей метод хешування у свій код, а потім виведіть результат у консоль. Знову ж таки, змінні вже визначені на сервері, тому вам не потрібно їх змінювати. Ви можете помітити, що хоча ви хешуєте той самий пароль, що й у асинхронній функції, результат у консолі відрізняється — це через те, що сіль генерується випадковим чином щоразу, про що свідчать перші 22 символи у третьому рядку хешу. Тепер, щоб порівняти введений пароль із новим синхронним хешем, ви використовуєте метод compareSync:

```js
var result = bcrypt.compareSync(myPlaintextPassword, hash);
```

Результатом буде булеве значення true або false.

# --instructions--

Додайте функцію і виведіть результат у консоль, щоб побачити, як вона працює.

Відправте свою сторінку, коли вважатимете, що все зроблено правильно.

# --hints--

Синхронний хеш має бути згенерований і правильно порівняний.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_SYNC[^]*hash.*=.*bcrypt.hashSync.*myPlaintextPassword( |),( |)saltRounds[^]*END_SYNC/gi,
  'You should call bcrypt.hashSync on myPlaintextPassword with saltRounds'
);
assert.match(
  data,
  /START_SYNC[^]*result.*=.*bcrypt.compareSync.*myPlaintextPassword( |),( |)hash[^]*END_SYNC/gi,
  'You should call bcrypt.compareSync on myPlaintextPassword with the hash generated in the last line'
);
```
