---
id: 58a25bcef9fc0f352b528e7c
title: Зрозуміти хеші BCrypt
challengeType: 2
forumTopicId: 301586
dashedName: understand-bcrypt-hashes
---

# --description--

Для наступних завдань ви працюватимете з новим стартовим проєктом, який відрізняється від попереднього. Ви можете знайти новий стартовий проєкт для клонування на <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Хеші BCrypt дуже безпечні. Хеш — це фактично відбиток оригінальних даних — завжди унікальний. Це досягається шляхом подачі оригінальних даних у алгоритм і повернення результату фіксованої довжини. Щоб ускладнити цей процес і зробити його більш безпечним, ви також можете *солити* свій хеш. Соління хешу означає додавання випадкових даних до оригінальних перед процесом хешування, що робить хеш ще складнішим для злому.

Хеші BCrypt завжди виглядають як `$2a$13$ZyprE5MRw2Q3WpNOGZWGbeG7ADUre1Q8QO.uUUtcbqloU0yvzavOm`, які мають певну структуру. Перша невелика частина даних `$2a` визначає, який тип алгоритму хешування було використано. Наступна частина `$13` визначає *вартість*. Вартість — це скільки ресурсів потрібно для обчислення хешу. Вона вимірюється за логарифмічною шкалою 2^вартість і визначає, скільки разів дані пропускаються через алгоритм хешування. Наприклад, при вартості 10 ви можете хешувати 10 паролів за секунду на середньому комп’ютері, але при вартості 15 це займає 3 секунди на один хеш... і ще далі, при вартості 31 це займе кілька днів для обчислення одного хешу. Вартість 12 вважається дуже безпечною на цей час. Остання частина вашого хешу `$ZyprE5MRw2Q3WpNOGZWGbeG7ADUre1Q8QO.uUUtcbqloU0yvzavOm` виглядає як великий рядок чисел, крапок і літер, але насправді це дві окремі частини інформації. Перші 22 символи — це сіль у відкритому вигляді, а решта — захешований пароль!

# --instructions--

Додайте весь свій код для цих уроків у файл `server.js` між кодом, з якого ми почали. Не змінюйте і не видаляйте код, який ми додали для вас.

BCrypt вже додано як залежність, тому підключіть його як `bcrypt` у вашому сервері.

Відправте свою сторінку, коли вважатимете, що зробили все правильно.

# --hints--

BCrypt має бути залежністю.

```js
const response = await fetch(code + '/_api/package.json');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
var packJson = JSON.parse(data);
assert.property(
  packJson.dependencies,
  'bcrypt',
  'Your project should list "bcrypt" as a dependency'
);
```

BCrypt має бути підключений правильно.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /bcrypt.*=.*require.*('|")bcrypt('|")/gi,
  'You should correctly require and instantiate bcrypt as bcrypt.'
);
```
