---
id: 58a25bcff9fc0f352b528e7d
title: Хешувати та порівнювати паролі асинхронно
challengeType: 2
forumTopicId: 301578
dashedName: hash-and-compare-passwords-asynchronously
---

# --description--

Нагадаємо, цей проєкт створюється на основі стартового проєкту, скопійованого з <a href="https://github.com/freeCodeCamp/boilerplate-bcrypt/" target="_blank" rel="noopener noreferrer nofollow">GitHub</a>.

Оскільки хешування розраховане на високе навантаження процесора, рекомендується виконувати його асинхронно на сервері, щоб уникнути блокування вхідних з’єднань під час хешування. Все, що потрібно зробити, щоб асинхронно хешувати пароль — це викликати

```js
bcrypt.hash(myPlaintextPassword, saltRounds, (err, hash) => {
  /*Store hash in your db*/
});
```

# --instructions--

Додайте цю функцію хешування на свій сервер (ми вже визначили змінні, які використовуються у функції, щоб ви могли їх використати) і виведіть результат у консоль, щоб побачити! Зазвичай на цьому етапі ви б зберегли хеш у базі даних.

Тепер, коли потрібно перевірити, чи співпадають нові вхідні дані з хешем, просто використовуйте функцію порівняння.

```js
bcrypt.compare(myPlaintextPassword, hash, (err, res) => {
  /*res == true or false*/
});
```

Додайте це у вашу існуючу функцію хешування (оскільки потрібно дочекатися завершення хешування перед викликом функції порівняння) після того, як ви вивели готовий хеш, і виведіть 'res' у консоль у функції порівняння. Ви повинні побачити у консолі хеш, а потім виведеться 'true'! Якщо змінити 'myPlaintextPassword' у функції порівняння на 'someOtherPlaintextPassword', тоді має вивестися хиба.

```js
bcrypt.hash('passw0rd!', 13, (err, hash) => {
  console.log(hash);
  //$2a$12$Y.PHPE15wR25qrrtgGkiYe2sXo98cjuMCG1YwSI5rJW1DSJp0gEYS
  bcrypt.compare('passw0rd!', hash, (err, res) => {
    console.log(res); //true
  });
});

```

Відправте свою сторінку, коли вважатимете, що зробили все правильно.

# --hints--

Асинхронний хеш має бути створений і правильно порівняний.

```js
const response = await fetch(code + '/_api/server.js');
if (!response.ok) {
  throw Error(await response.text());
}
const data = await response.text();
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash.*myPlaintextPassword( |),( |)saltRounds( |),( |).*err( |),( |)hash[^]*END_ASYNC/gi,
  'You should call bcrypt.hash on myPlaintextPassword and saltRounds and handle err and hash as a result in the callback'
);
assert.match(
  data,
  /START_ASYNC[^]*bcrypt.hash[^]*bcrypt.compare.*myPlaintextPassword( |),( |)hash( |),( |).*err( |),( |)res[^]*}[^]*}[^]*END_ASYNC/gi,
  'Nested within the hash function should be the compare function comparing myPlaintextPassword to hash'
);
```
