---
id: 6823f9df49cc206af5471a30
title: Налагодьте профіль Кемпербота
challengeType: 25
dashedName: lab-debug-camperbots-profile-page
---

# --description--

Кемпербот намагається налаштувати свій профіль. Він попросив друзів глянути на код і вони знайшли декілька помилок.

Ваше завдання — розв’язати помилки, щоб Кемпербот продовжив налаштовувати свій профіль. Виконайте історію користувача нижче і натисніть на «Запустити тест», щоб перевірити, чи всі помилки виправлено.

**Історія користувача:**

1. Кемпербот намагається скористатись елементом `heading2`, але такого елемента не існує. Виправте ці теги на такі, які потрібні для заголовка другого рівня.
2. Кемпербот намагається додати два абзаци за допомогою `pp`, але такі елементи теж не існують. Виправте їх на правильні теги для елементів-абзаців.
3. Кемпербот використовує елемент `h3` для підзаголовка `Background and Interests`, але має синтаксичну помилку. Знайдіть і розв’яжіть її.

# --hints--

На сторінці має бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Не видаляйте елемент `<h1>Hello from Camperbot!</h1>` зі сторінки.

```js
assert.match(code, /\<h1\s*\>\s*Hello\s*from\s*Camperbot\s*!\s*\<\/h1\s*\>/i);
```

На сторінці має бути лише один елемент `h2`.

```js
assert.lengthOf(document.querySelectorAll("h2"), 1);
```

Ви повинні мати елемент `h2` з текстом `About`. Ось приклад: `<h2>Тут текст</h2>`.

```js
assert.match(code, /\<h2\s*\>\s*About\s*\<\/h2\s*\>/i);
```

Ви повинні мати елемент-абзац з текстом `My name is Camperbot and I love learning new things.` Ось приклад: `<p>Тут текст</p>`.

```js
assert.match(code, /\<p\s*\>\s*My\s*name\s*is\s*camperbot\s*and\s*I\s*love\s*learning\s*new\s*things\s*\.\s*\<\/p\s*\>/i);
```

Ви повинні мати елемент-абзац з текстом `I enjoy solving puzzles.` Ось приклад: `<p>Тут текст</p>`.

```js
assert.match(code, /\<p\s*\>\s*I\s*enjoy\s*solving\s*puzzles\s*\.\s*\<\/p\s*\>/i);
```

На сторінці має бути два елементи-абзаци.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

На сторінці має бути лише один елемент `h3`.

```js
assert.lengthOf(document.querySelectorAll("h3"), 1);
```

Ви повинні мати елемент `h3` з текстом `Background and Interests`. Ось приклад: `<h3>Тут текст</h3>`.

```js
assert.match(code, /\<h3\s*\>\s*Background\s*and\s*Interests\s*\<\/h3\s*\>/i);
```

Загалом на сторінці має бути п’ять елементів.

```js
assert.lengthOf(document.querySelectorAll("body *:not(script)"), 5);
```

# --seed--

## --seed-contents--

```html
<h1>Hello from Camperbot!</h1>

<heading2>About</heading2>

<pp>My name is Camperbot and I love learning new things.</pp>

<h3>Background and Interests<h3/>
<pp>I enjoy solving puzzles.</pp>
```

# --solutions--

```html
<h1>Hello from Camperbot!</h1>

<h2>About</h2>

<p>My name is Camperbot and I love learning new things.</p>

<h3>Background and Interests</h3>
<p>I enjoy solving puzzles.</p>
```
