---
id: a6e40f1041b06c996f7b2406
title: 建置第一個元素尋找器
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

在本實驗中，你將創建一個函式（程式），該函式會檢查一個陣列並傳回第一個通過測試函式（也稱為「真值測試」）的元素。 

該函式會迭代陣列並使用提供的測試函式測試每個元素。最後，它會傳回測試函式返回 `true` 的第一個元素。

範例：

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `findElement` 的函式（程式），它接受一個陣列和一個函式（程式）作為引數。
2. 該函式應傳回陣列中第一個通過真值測試的項目。這表示，呼叫傳入的函式 `func`，給定一個元素 `x`，如果 `func(x)` 是 `true`，則表示通過真值測試。
3. 如果沒有元素通過測試，該函式（程式）應該傳回 `undefined`。

# --hints--

你應該有一個 `findElement` 函式（程式）。

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` 應該傳回 `8`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` 應該傳回 `undefined`。

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` 應該傳回 `3`。

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` 應該傳回 `"javascript"`。

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` 應該傳回 `undefined`。

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` 應該傳回 `2`。

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` 應該傳回 `undefined`。

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
