---
id: 66f53dc2c5bd6a11d6c3282f
title: 建置一個密碼產生器應用程式
challengeType: 26
dashedName: lab-password-generator
---

# --description--

在這個實驗中，你將透過建置一個隨機密碼產生器來練習使用函式（程式）。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `generatePassword` 的函式（程式），該函式帶有一個參數，用來表示產生密碼的長度。你可以將該參數命名為任何你喜歡的名稱。
2. 你的函式應該傳回一個字串，該字串表述一個隨機產生的密碼。你應該使用以下字串和不同的 `Math` 方法來幫助你傳回一個包含隨機字元的新字串：`ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`。
3. 你應該定義一個名為 `password` 的變數，並將呼叫 `generatePassword` 函式（程式）且帶有表述所需密碼長度的數值引數的結果指定給它。
4. 你應該有一個 `console.log`，它會記錄一個由字串 `Generated password:` 和 `password` 變數以空格分隔後串接而成的單一字串。

# --hints--

你應該有一個帶有參數的 `generatePassword` 函式。你可以將參數命名為任何你喜歡的名稱。

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

你的 `generatePassword` 函式應該傳回一個字串。

```js
const result = generatePassword(5);
assert.isString(result);
```

你的 `generatePassword` 函式應該傳回一個長度正確的新字串。

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

你的函式應該傳回一個包含有效字元的隨機產生密碼。

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

你的函式應該傳回一個隨機產生且包含多於一個獨特字元的密碼。

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

你的函式（程式）每次被呼叫時，應該傳回一個新的隨機字串。

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

你應該有一個 `password` 變數。

```js
assert.isDefined(password);
```

你的 `password` 變數應該是一個字串。

```js
assert.isString(password);
```

你應該以數字引數呼叫 `generatePassword` 函式，並將傳回的密碼儲存在 `password` 變數中。

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

你應該記錄一個單一字串，該字串結合了 `Generated password:` 和 `password`，兩者以單一空格分隔，並使用 `+` 或模板字面值。

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
