---
id: 587d824a367417b2b2512c44
title: Перевірка ціни акцій
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

Створіть повноцінний JavaScript-застосунок, функціонально схожий на цей: <a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>.

Оскільки всі надійні API для цін акцій вимагають ключ API, ми створили обхідне рішення. Використовуйте <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a>, щоб отримувати актуальну інформацію про ціни акцій без необхідності реєструвати власний ключ.

Робота над цим проєктом передбачає написання коду одним із таких способів:

-   Клонувати <a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій GitHub</a> і завершити проєкт локально.
-   Використати будь-який конструктор сайтів на ваш вибір для завершення проєкту. Обов’язково включіть усі файли з нашого репозиторію GitHub.

# --instructions--

1.  Встановіть змінну середовища `NODE_ENV` зі значенням `test`, без лапок
2.  Завершіть проєкт у `routes/api.js` або створивши обробник/контролер
3.  Ви додасте всі функції безпеки у `server.js`
4.  Ви створите всі функціональні тести у `tests/2_functional-tests.js`

**Примітка** Щодо конфіденційності: через вимогу приймати лише 1 лайк з однієї IP-адреси, вам доведеться зберігати IP-адреси. Важливо дотримуватися законів про захист даних, таких як Загальний регламент захисту даних. Один із варіантів — отримати дозвіл на збереження даних користувача, але значно простіше анонімізувати їх. Для цього завдання пам’ятайте анонімізувати IP-адреси перед збереженням у базі даних. Якщо потрібні ідеї, як це зробити, ви можете вибрати хешування даних, їх усічення або встановлення частини IP-адреси в 0.

Напишіть такі тести у `tests/2_functional-tests.js`:

-   Перегляд однієї акції: GET-запит до `/api/stock-prices/`
-   Перегляд однієї акції з лайком: GET-запит до `/api/stock-prices/`
-   Повторний перегляд тієї ж акції з лайком: GET-запит до `/api/stock-prices/`
-   Перегляд двох акцій: GET-запит до `/api/stock-prices/`
-   Перегляд двох акцій з лайками: GET-запит до `/api/stock-prices/`

# --hints--

Ви можете надати власний проєкт, а не приклад URL.

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

Ви повинні встановити політики безпеки вмісту так, щоб дозволити завантаження скриптів і CSS лише з вашого сервера.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

Ви можете надіслати `GET`-запит до `/api/stock-prices`, передаючи символ акції NASDAQ у параметрі запиту `stock`. Повернутий об’єкт міститиме властивість з назвою `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

Властивість `stockData` включає символ `stock` як рядок, `price` як число та `likes` як число.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

Ви також можете передати поле `like` зі значенням `true` (булевий), щоб додати свій лайк до акції(й). Приймається лише 1 лайк з однієї IP-адреси.

```js

```

Якщо передати 2 акції, повернене значення буде масивом з інформацією про обидві акції. Замість `likes` воно відобразить `rel_likes` (різницю між лайками обох акцій) для обох об’єктів `stockData`.

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

Всі 5 функціональних тестів завершені та проходять успішно.

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```
