---
id: 5e601c775ac9d0ecd8b94aff
title: Безпечна багатокористувацька гра в реальному часі
challengeType: 4
forumTopicId: 462375
dashedName: secure-real-time-multiplayer-game
---

# --description--

Розробіть 2D багатокористувацьку гру в реальному часі, використовуючи HTML Canvas API та Socket.io, яка функціонально схожа на цю: <a href="https://secure-real-time-multiplayer-game.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://secure-real-time-multiplayer-game.freecodecamp.rocks/</a>. Робота над цим проєктом передбачає написання коду одним із таких способів:

-   Клонуйте <a href="https://github.com/freeCodeCamp/boilerplate-project-secure-real-time-multiplayer-game/" target="_blank" rel="noopener noreferrer nofollow">цей репозиторій GitHub</a> і завершіть проєкт локально.
-   Використайте будь-який конструктор сайтів на свій вибір для завершення проєкту. Обов’язково включіть усі файли з нашого репозиторію GitHub.

# --instructions--

Створіть безпечну багатокористувацьку гру, у якій кожен гравець може рухати свого аватара, є принаймні один збиральний елемент, а рейтинг гравців обчислюється на основі їхнього рахунку.

Для деталей зверніться до тестів нижче.

Переконайтеся, що ваша гра безпечна! Включіть такі заходи безпеки:

- Клієнт не повинен мати змогу вгадувати або визначати MIME-тип
- Запобігайте XSS-атакам
- Не кешуйте нічого з сайту на клієнті
- Заголовки повинні вказувати, що сайт працює на `PHP 7.4.3`

**Примітка**: `helmet@^3.21.3` потрібен для історій користувача. Це означає, що вам доведеться використовувати попередню версію документації Helmet, щоб дізнатися, як реалізувати історії користувача.

# --hints--

Ви можете надати свій власний проєкт, а не приклад URL.

```js
  assert(
    !/.*\/secure-real-time-multiplayer-game\.freecodecamp\.rocks/.test(
      code
    )
  );
```

Кілька гравців можуть підключатися до сервера і грати.

```js

```

Кожен гравець має аватар.

```js

```

Кожен гравець представлений об’єктом, створеним класом `Player` у `Player.mjs`.

```js

```

Щонайменше, кожен об’єкт гравця повинен містити унікальний `id`, `score`, а також координати `x` і `y`, що представляють поточне положення гравця.

```js

```

У грі є принаймні один тип збирального елемента. Завершіть клас `Collectible` у `Collectible.mjs`, щоб реалізувати це.

```js

```

Щонайменше, кожен об’єкт збирального елемента, створений класом `Collectible`, повинен містити унікальний `id`, `value`, а також координати `x` і `y`, що представляють поточне положення елемента.

```js

```

Гравці можуть рухати свого аватара за допомогою клавіш WASD і/або стрілок. Завершіть метод `movePlayer` у `Player.mjs`, щоб реалізувати це.

```js

```

Метод `movePlayer` повинен приймати два аргументи: рядок "up", "down", "left" або "right" і число, що визначає кількість пікселів, на яку має змінитися позиція гравця. `movePlayer` має коригувати координати `x` і `y` об’єкта гравця, з якого він викликається.

```js

```

Рахунок гравця має використовуватися для обчислення його рейтингу серед інших гравців. Завершіть метод `calculateRank` у класі `Player`, щоб реалізувати це.

```js

```

Метод `calculateRank` повинен приймати масив об’єктів, що представляють усіх підключених гравців, і повертати рядок `Rank: currentRanking/totalPlayers`. Наприклад, у грі з двома гравцями, якщо у Гравця A рахунок 3, а у Гравця B — 5, `calculateRank` для Гравця A має повернути `Rank: 2/2`.

```js

```

Гравці можуть стикатися зі збиральним елементом. Завершіть метод `collision` у `Player.mjs`, щоб реалізувати це.

```js

```

Метод `collision` повинен приймати об’єкт збирального елемента як аргумент. Якщо аватар гравця перетинається з елементом, метод `collision` має повернути `true`.

```js

```

Усі гравці синхронізовані.

```js

```

Гравці можуть відключатися від гри у будь-який момент.

```js

```

Запобігайте спробам клієнта вгадувати або визначати MIME-тип.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-content-type-options'], 'nosniff');
```

Запобігайте атакам міжсайтового скриптингу (XSS).

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-xss-protection'], '1; mode=block');
```

Нічого з сайту не кешується на клієнті.

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['surrogate-control'], 'no-store');
  assert.equal(
    parsed.headers['cache-control'],
    'no-store, no-cache, must-revalidate, proxy-revalidate'
  );
  assert.equal(parsed.headers['pragma'], 'no-cache');
  assert.equal(parsed.headers['expires'], '0');
```

Заголовки вказують, що сайт працює на "PHP 7.4.3", хоча це не так (з міркувань безпеки).

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.equal(parsed.headers['x-powered-by'], 'PHP 7.4.3');
```
