---
id: 614385513d91ae5c251c2052
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop erstellen Sie eine Magazinseite anhand einer Pressemitteilung aus einer früheren Version des freeCodeCamp-Curriculums.

Beginnen Sie damit, ein `link`-Element für das `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`-Schriftart-Stylesheet, ein `link`-Element für das `https://use.fontawesome.com/releases/v5.8.2/css/all.css`-FontAwesome-Stylesheet und ein `link`-Element für Ihr `./styles.css`-Stylesheet hinzuzufügen.

Ihr Schriftart-Stylesheet lädt drei verschiedene Schriftarten: `Anton`, `Baskervville` und `Raleway`.

# --hints--

Ihr Code sollte drei `link`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Ihre `link`-Elemente sollten sich innerhalb Ihres `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

Ihre drei `link`-Elemente sollten ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines Ihrer Link-Elemente sollte das `href`-Attribut auf `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` gesetzt haben.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

Eines Ihrer Link-Elemente sollte das `href`-Attribut auf `https://use.fontawesome.com/releases/v5.8.2/css/all.css` gesetzt haben.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

Eines Ihrer `link`-Elemente sollte das `href`-Attribut auf `styles.css` gesetzt haben.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
