---
id: 6143862a5e54455d262c212e
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie innerhalb Ihres `body` ein `main`-Element. Erstellen Sie dann in diesem Element ein `section` mit einem `class`, das auf `heading` gesetzt ist.

# --hints--

Sie sollten ein `main`-Element haben.

```js
assert.exists(document.querySelector('main'));
```

Ihr `main`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert.equal(document.querySelector('main')?.parentElement?.localName, 'body');
```

Sie sollten ein `section`-Element haben.

```js
assert.exists(document.querySelector('section'));
```

Ihr `section`-Element sollte sich innerhalb Ihres `main`-Elements befinden.

```js
assert.equal(document.querySelector('section')?.parentElement?.localName, 'main');
```

Ihr `section`-Element sollte das `class` auf `heading` gesetzt haben.

```js
assert.isTrue(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
