---
id: 6143869bb45bd85e3b1926aa
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie innerhalb Ihres `.heading`-Elements ein `header`-Element mit dem `class`, das auf `hero` gesetzt ist.

Erstellen Sie in diesem Element ein `img`-Element mit dem `src`, das auf `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` gesetzt ist, dem `alt`, das auf `freecodecamp logo` gesetzt ist, und dem `class`, das auf `hero-img` gesetzt ist.

Das `loading`-Attribut eines `img`-Elements kann auf `lazy` gesetzt werden, um dem Browser mitzuteilen, die Bildressource erst dann zu laden, wenn sie benötigt wird (also wenn der Benutzer das Bild in den sichtbaren Bereich scrollt). Als zusätzlicher Vorteil werden lazy geladene Elemente erst geladen, nachdem die nicht-lazy Elemente geladen sind – das bedeutet, dass Nutzer mit langsamen Internetverbindungen den Inhalt Ihrer Seite sehen können, ohne auf das Laden der Bilder warten zu müssen.

Geben Sie Ihrem neuen `img`-Element ein `loading`-Attribut, das auf `lazy` gesetzt ist.

# --hints--

Sie sollten ein `header`-Element erstellen.

```js
assert.exists(document.querySelector('header'));
```

Ihr `header`-Element sollte sich innerhalb Ihres `.heading`-Elements befinden.

```js
assert.equal(document.querySelector('header')?.parentElement?.className, 'heading');
```

Ihr `header`-Element sollte das `class` auf `hero` gesetzt haben.

```js
assert.equal(document.querySelector('header')?.className, 'hero');
```

Sie sollten ein `img`-Element erstellen.

```js
assert.exists(document.querySelector('img'));
```

Ihr `img`-Element sollte sich innerhalb Ihres `header`-Elements befinden.

```js
assert.equal(document.querySelector('img')?.parentElement?.localName, 'header');
```

Ihr `img`-Element sollte das `src` auf `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png` gesetzt haben.

```js
assert.equal(document.querySelector('img')?.getAttribute('src'), 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

Ihr `img`-Element sollte das `alt` auf `freecodecamp logo` gesetzt haben.

```js
assert.equal(document.querySelector('img')?.getAttribute('alt'), 'freecodecamp logo');
```

Ihr `img`-Element sollte das `loading`-Attribut auf `lazy` gesetzt haben.

```js
assert.equal(document.querySelector('img')?.getAttribute('loading'), 'lazy');
```

Ihr `img`-Element sollte das `class` auf `hero-img` gesetzt haben.

```js
assert.equal(document.querySelector('img')?.className, 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
