---
id: 61438ec09438696391076d6a
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Unter Ihrem `.heading`-Element erstellen Sie ein neues `section`-Element mit dem `class`, das auf `text` gesetzt ist. Darin erstellen Sie ein `p`-Element mit dem `class`, das auf `first-paragraph` gesetzt ist, und dem folgenden Text:

```markup
Soon the freeCodeCamp curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you'll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.
```

# --hints--

Sie sollten ein neues `section`-Element erstellen.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Ihr neues `section`-Element sollte nach Ihrem `.heading`-Element kommen.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.previousElementSibling?.className, 'heading');
```

Ihr neues `section`-Element sollte das `class` auf `text` gesetzt haben.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.className, 'text');
```

Sie sollten ein neues `p`-Element innerhalb Ihres `.text`-Elements erstellen.

```js
assert.lengthOf(document.querySelector('.text')?.querySelectorAll('p'), 1);
```

Ihr neues `p`-Element sollte das `class` auf `first-paragraph` gesetzt haben.

```js
assert.equal(document.querySelector('.text p')?.className, 'first-paragraph');
```

Ihr neues `p`-Element sollte den bereitgestellten Text enthalten.

```js
assert.equal(document.querySelector('.first-paragraph')?.innerText, 'Soon the freeCodeCamp curriculum will be 100% project-driven learning. Instead of a series of coding challenges, you\'ll learn through building projects - step by step. Before we get into the details, let me emphasize: we are not changing the certifications. All 6 certifications will still have the same 5 required projects. We are only changing the optional coding challenges.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
            <a href="https://freecodecamp.org" target="_blank" rel="noreferrer"
              >freeCodeCamp</a
            >
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
        <div class="social-icons">
          <a href="https://www.facebook.com/freecodecamp/">
            <i class="fab fa-facebook-f"></i>
          </a>
          <a href="https://twitter.com/freecodecamp/">
            <i class="fab fa-twitter"></i>
          </a>
          <a href="https://instagram.com/freecodecamp">
            <i class="fab fa-instagram"></i>
          </a>
          <a href="https://www.linkedin.com/school/free-code-camp/">
            <i class="fab fa-linkedin-in"></i>
          </a>
          <a href="https://www.youtube.com/freecodecamp">
            <i class="fab fa-youtube"></i>
          </a>
        </div>
      </section>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
