---
id: 614387cbefeeba5f3654a291
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Nach Ihrem `header`-Element erstellen Sie ein `div` mit dem `class`, das auf `author` gesetzt ist.

Innerhalb dieses `div` erstellen Sie ein `p`-Element mit dem `class`, das auf `author-name` gesetzt ist, und geben ihm den Text `By freeCodeCamp`. Wickeln Sie den `freeCodeCamp`-Abschnitt in ein `a`-Element mit dem `href`, das auf `https://freecodecamp.org` gesetzt ist, und dem `target`, das auf `_blank` gesetzt ist.


Darunter fügen Sie ein zweites `p`-Element mit der Klasse `publish-date` und dem Text `March 7, 2019` hinzu.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert.exists(document.querySelector('div'));
```

Ihr `div`-Element sollte nach Ihrem `header`-Element kommen.

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

Ihr `div`-Element sollte das `class` auf `author` gesetzt haben.

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

Sie sollten zwei neue `p`-Elemente erstellen.

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

Ihre zwei neuen `p`-Elemente sollten innerhalb Ihres `div`-Elements sein.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Ihr erstes neues `p`-Element sollte ein `class` haben, das auf `author-name` gesetzt ist.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

Ihr erstes neues `p`-Element sollte den Text `By freeCodeCamp` haben.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

Ihr zweites neues `p`-Element sollte ein `class` haben, das auf `publish-date` gesetzt ist.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

Ihr zweites neues `p`-Element sollte den Text `March 7, 2019` haben.

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

Sie sollten ein neues `a`-Element erstellen.

```js
assert.exists(document.querySelector('a'));
```

Ihr `a`-Element sollte innerhalb Ihres ersten neuen `p`-Elements sein.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

Ihr `a`-Element sollte das `href` auf `https://freecodecamp.org` gesetzt haben.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

Ihr `a`-Element sollte das `target` auf `_blank` gesetzt haben.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

Ihr `a`-Element sollte den Text `freeCodeCamp` umschließen.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
