---
id: 6165d3b702a5d92ad970b30c
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie nach Ihrem `img`-Element ein `h1`-Element mit dem `class`, das auf `hero-title` gesetzt ist, und dem Text, der auf `OUR NEW CURRICULUM` gesetzt ist, hinzu, gefolgt von einem `p`-Element mit dem `class`, das auf `hero-subtitle` gesetzt ist, und dem Text, der auf `Our efforts to restructure our curriculum with a more project-based focus` gesetzt ist.

# --hints--

Sie sollten ein `h1`-Element erstellen.

```js
assert.exists(document.querySelector('h1'));
```

Ihr `h1`-Element sollte nach Ihrem `img`-Element kommen.

```js
assert.equal(document.querySelector('h1')?.previousElementSibling?.localName, 'img');
```

Ihr `h1`-Element sollte das `class` auf `hero-title` gesetzt haben.

```js
assert.equal(document.querySelector('h1')?.className, 'hero-title');
```

Ihr `h1`-Element sollte den Text auf `OUR NEW CURRICULUM` gesetzt haben.

```js
assert.equal(document.querySelector('h1')?.textContent, 'OUR NEW CURRICULUM');
```

Sie sollten ein neues `p`-Element erstellen.

```js
assert.exists(document.querySelector('p'));
```

Ihr `p`-Element sollte nach Ihrem `h1`-Element kommen.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ihr `p`-Element sollte das `class` auf `hero-subtitle` gesetzt haben.

```js
assert.equal(document.querySelector('p')?.className, 'hero-subtitle');
```

Ihr `p`-Element sollte den Text auf `Our efforts to restructure our curriculum with a more project-based focus` gesetzt haben.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </header>
      </section>
    </main>
  </body>
</html>
```

```css

```
