---
id: 6169cd8a558aa8434e0ad7f6
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Der HTTP-Kopfzeile `Referer` enthält Informationen über die Adresse oder URL einer Seite, von der ein Benutzer möglicherweise kommt. Diese Informationen können in der Analyse verwendet werden, um beispielsweise zu verfolgen, wie viele Benutzer von Ihrer Seite freecodecamp.org besuchen. Wenn Sie das `rel`-Attribut auf `noreferrer` setzen, wird diese Information aus der HTTP-Anfrage ausgelassen. Geben Sie Ihrem `a`-Element ein `rel`-Attribut mit dem Wert `noreferrer`.

# --hints--

Ihr `a`-Element sollte das `rel` auf `noreferrer` gesetzt haben.

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('rel'), 'noreferrer');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
        <div class="author">
          <p class="author-name">
            By
--fcc-editable-region--
            <a href="https://freecodecamp.org" target="_blank"
              >freeCodeCamp</a
            >
--fcc-editable-region--
          </p>
          <p class="publish-date">March 7, 2019</p>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
