---
id: 66f17db06803d11a1bd19a20
title: 퀴즈 게임 만들기
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `questions`라는 배열을 만들어야 합니다.
1. `questions` 배열에는 `category`, `question`, `choices`, `answer` 키를 가진 객체가 최소 다섯 개 있어야 합니다.
1. `category` 키는 질문 카테고리를 나타내는 문자열 값을 가져야 합니다.
1. `question` 키는 질문을 나타내는 문자열 값을 가져야 합니다.
1. `choices` 키는 질문에 대한 세 가지 대체 답변을 포함하는 문자열 배열 값을 가져야 합니다.
1. `answer` 키는 질문에 대한 정답을 나타내는 문자열 값을 가져야 합니다. 또한 `answer` 값은 `choices` 배열에 포함되어야 합니다.
1. `getRandomQuestion`라는 이름의 함수를 만들어야 하며, 이 함수는 질문 배열을 매개변수로 받아 배열에서 무작위 질문 객체를 반환해야 합니다.
1. `getRandomComputerChoice`라는 이름의 함수를 만들어야 하며, 이 함수는 사용 가능한 선택지 배열을 매개변수로 받아 선택된 질문에 대한 무작위 답변을 반환해야 합니다.
1. `getResults`라는 이름의 함수를 만들어야 하며, 이 함수는 첫 번째 매개변수로 질문 객체를, 두 번째 매개변수로 컴퓨터의 선택을 받습니다. 답변이 맞으면 `The computer's choice is correct!`를 반환하고, 그렇지 않으면 `The computer's choice is wrong. The correct answer is: <correct-answer>`를 반환하는데, 여기서 `<correct-answer>`는 선택된 질문의 정답 값입니다.

# --hints--

`questions`라는 배열을 만들어야 합니다.

```js
assert.isArray(questions);
```

`questions` 배열에는 `category`, `question`, `choices`, `answer` 키를 가진 객체가 최소 다섯 개 있어야 합니다.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

`category` 키는 질문 카테고리를 나타내는 문자열 값을 가져야 합니다.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

`question` 키는 질문을 나타내는 문자열 값을 가져야 합니다.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

`choices` 키는 서로 다른 세 개의 문자열을 포함하는 배열 값을 가져야 합니다.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

`answer` 키는 문자열 값을 가져야 합니다.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

`answer` 값은 `choices` 배열에 포함되어야 합니다.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

`getRandomQuestion`라는 이름의 함수를 만들어야 하며, 이 함수는 질문 배열을 매개변수로 받아 배열에서 무작위 질문 객체를 반환해야 합니다.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

`getRandomComputerChoice`라는 이름의 함수를 만들어야 하며, 이 함수는 사용 가능한 선택지 배열을 매개변수로 받아 선택된 질문에 대한 무작위 답변을 반환해야 합니다.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

`getResults`라는 이름의 함수를 만들어야 합니다.

```js
assert.isFunction(getResults);
```

`getResults` 함수는 첫 번째 매개변수로 질문 객체를, 두 번째 매개변수로 컴퓨터의 선택을 받아야 합니다.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

컴퓨터의 선택이 정답과 일치하면 `getResults`는 `The computer's choice is correct!`를 반환해야 합니다.

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

컴퓨터의 선택이 정답과 일치하지 않으면 `getResults`는 `The computer's choice is wrong. The correct answer is: <correct-answer>`를 반환해야 하며, 여기서 `<correct-answer>`는 선택된 질문의 정답 값입니다.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

`getResults` 함수는 부분 문자열 비교가 아닌 정확한 동등성 비교를 사용해야 합니다.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
