---
id: 65f0694fb2296f3caadf8347
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

La lunghezza di un vettore $\mathbf{a}$, o norma, è solitamente indicata come $\\| \mathbf{a} \\|$. Può essere calcolata come la radice quadrata della somma dei suoi componenti al quadrato:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

Calcola la norma del vettore e restituisci il risultato dal tuo metodo `norm`.

# --hints--

Il tuo metodo `norm` dovrebbe usare la formula fornita per calcolare e restituire un numero che rappresenta la norma dell'istanza vettore corrente. Non approssimare il valore.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
