---
id: 65f03d9f92eac9183a4d3281
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Un vettore può essere definito da due coordinate, `x` e `y`, nel piano euclideo. La distanza tra l'origine degli assi e il punto `(x, y)` sarà la sua lunghezza, o norma. E la direzione del vettore punterà verso `(x, y)`.

<img alt="un vettore bidimensionale con coordinate (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

All'interno della classe `Vector`, crea un metodo `__init__` e dagli tre parametri, `self`, `x` e `y`.

# --hints--

Dovresti definire un metodo `__init__` all'interno della classe `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Il tuo metodo `__init__` dovrebbe prendere tre parametri: `self`, `x` e `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
