---
id: 65f055b9190fc41ca35549b8
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Python offre vari metodi che includono sia un doppio underscore iniziale sia finale nei loro nomi. Potresti già conoscere alcuni, come `__init__` e `__str__`. Questi metodi, che seguono il pattern di denominazione `__<name>__`, sono chiamati metodi speciali, metodi magici o metodi *dunder* (che sta per double underscore).

Definire metodi speciali in una classe modifica il comportamento di quella classe. Vengono chiamati "sotto il cofano" in situazioni specifiche (ad esempio `__init__` durante l'istanza, `__str__` quando l'oggetto viene stampato o passato a `str()`). In questo progetto imparerai alcuni dei metodi speciali più comunemente usati.

Per ora, assegna `x` all'attributo `x` dell'oggetto `Vector`.

# --hints--

Dovresti assegnare `x` a `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
