---
id: 65f058e06f34fd1f0ee6e55d
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Successivamente, all'interno della classe `Vector`, crea un metodo `norm` vuoto e assegnagli un parametro `self`.

# --hints--

La tua classe `Vector` dovrebbe avere un metodo `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Il tuo metodo `norm` dovrebbe avere un parametro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
