---
id: 65f06a8e5a57673d700c79c3
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

All'esterno della classe `Vector`, crea un'istanza di `Vector` passando i numeri interi `2` e `3` come valori `x` e `y` e assegnala a una variabile chiamata `v1`.

# --hints--

Dovresti dichiarare una variabile `v1` e assegnarle `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
