---
id: 65f01969115f933073b6be03
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Un vettore è un oggetto che ha una lunghezza (o grandezza) e una direzione e non può essere espresso da un singolo numero. In fisica, i vettori sono comunemente usati per rappresentare forze, velocità, accelerazioni e altre grandezze.

In questo progetto costruirai uno spazio vettoriale, un insieme in cui è definita una serie di operazioni tra gli elementi di quell’insieme. Imparerai tutti i dettagli molto presto.

Per ora, inizia il progetto dichiarando una classe vuota chiamata `Vector`.

# --hints--

Dovresti dichiarare una classe vuota chiamata `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
