---
id: 65f3fa097f26b510db6c710b
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Ottimo, `norm` funziona come previsto. Ora, se provi a stampare `v1`, otterrai la rappresentazione stringa predefinita di un oggetto (qualcosa come `<__main__.Vector object at 0x11eb778>`).

All'interno della classe `Vector`, dichiara un metodo `__str__` vuoto per implementare una rappresentazione stringa leggibile. Ricorda di dargli un parametro `self`.

Fai attenzione a non stampare `v1` finché `__str__` non restituisce una stringa, altrimenti otterrai un `TypeError`, perché `__str__` deve sempre restituire una stringa.

# --hints--

Dovresti definire un metodo `__str__` all'interno della classe `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

Il tuo metodo `__str__` dovrebbe avere un parametro `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
