---
id: 65f3fa60a93b84110b3f1708
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

Vuoi restituire una rappresentazione stringa significativa. Per esempio, nel caso di `v1`, vuoi restituire una stringa contenente la tupla con i valori dei componenti del vettore: `(2, 3)`.

Dal metodo `__str__`, restituisci una stringa che rappresenti il vettore come una tupla contenente i componenti del vettore in ordine. Poi, esci dalla classe `Vector` e stampa `v1` per verificare il risultato.

# --hints--

Il tuo metodo `__str__` dovrebbe restituire una stringa che rappresenti il vettore come una tupla contenente i componenti del vettore in ordine.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Dovresti stampare `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
