---
id: 65f3fe07cc763212efe91285
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Un vettore può avere un numero `n` di dimensioni (componenti). Ecco una rappresentazione di un vettore tridimensionale:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Finora hai creato un vettore bidimensionale. Vuoi poter rappresentare vettori con un numero diverso di dimensioni senza riscrivere il codice necessario per ogni caso specifico. Per questo userai l'ereditarietà.

Inizia rinominando la classe `Vector` in `R2Vector` per specificare che questa classe gestirà vettori bidimensionali. Ricordati di modificare anche l'istanza di `v1`.

# --hints--

Dovresti rinominare la classe `Vector` in `R2Vector`. Ricordati di modificare anche l'istanza di `v1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
