---
id: 65f4535bbdb28d436ff3ddc9
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Come puoi vedere dall'output, `__dict__` contiene i valori dei tuoi attributi di istanza. Invece di aggiungere esplicitamente i quadrati di `self.x` e `self.y`, itererai sui valori archiviati in `__dict__` per calcolare il valore della norma.

All'interno del corpo del metodo `norm`, sostituisci il contenuto delle parentesi tonde con un'espressione generatrice che eleva ogni valore `val` in `self.__dict__.values()` alla potenza di `2`. Inoltre, passa quell'espressione generatrice come argomento alla funzione `sum`, in modo che tutti i quadrati vengano sommati prima di calcolare la radice quadrata.

# --hints--

Dovresti restituire `sum(val**2 for val in self.__dict__.values())**0.5` dal metodo `norm()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
