---
id: 65f40051d6b09a139f253e8e
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

L'ereditarietà ti permette di definire una classe a partire da una esistente. La nuova classe, chiamata figlio, eredita tutti i metodi e le proprietà della classe esistente, chiamata genitore.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

Nell'esempio sopra, la classe figlio `Oak` eredita da `Tree` e eredita il metodo `sprout` dalla classe genitore `Tree`.

Crea una nuova classe chiamata `R3Vector` e segui l'esempio sopra per farla ereditare dalla classe `R2Vector`.

# --hints--

Dovresti definire una nuova classe `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

La tua nuova classe `R3Vector` dovrebbe ereditare da `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
