---
id: 65f412c208c3791fee305acf
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

In Python, puoi imporre l’uso di argomenti keyword-only aggiungendo un `*` come argomento aggiuntivo alla firma della funzione o del metodo. Modifica entrambi i metodi `__init__` aggiungendo un `*` come secondo parametro (dopo `self`). Ogni parametro posizionato dopo richiederà l’uso di un argomento keyword nella chiamata della funzione/metodo.

Questo significa che devi modificare anche la chiamata a `super().__init__(x, y)`. Fallo assegnando a `x` il valore `x` e a `y` il valore `y`.

Infine, modifica l’istanza di `v1` e `v2` usando argomenti keyword.

# --hints--

Dovresti imporre gli argomenti keyword aggiungendo un `*` come secondo parametro in entrambi i tuoi metodi `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Dovresti modificare la chiamata a `super.__init__(x, y)` per usare argomenti keyword.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Dovresti modificare l’assegnazione di `v1` in modo che sia `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Dovresti modificare l’assegnazione di `v2` in modo che sia `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
