---
id: 65f407ea37ad6e181b90462e
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

Potresti assegnare ogni parametro `i` a `self.i` come hai fatto prima. Anche se in questo caso hai poche righe da ripetere, un modo per evitare questa ripetizione è usare la funzione `super()`. `super()` ti permette di riferirti implicitamente alla classe genitore: `super().__init__(x, y)` chiama il metodo `__init__` della classe genitore.

Aggiungi una chiamata `super().__init__(x, y)` al tuo metodo `__init__`.

# --hints--

Dovresti avere una chiamata `super().__init__(x, y)` nel tuo metodo `__init__`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
