---
id: 65f4520e363e2642f8112e33
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

Come puoi vedere, qualcosa non va. I metodi `norm` e `__str__` ereditati da `R2Vector` non possono adattarsi a un vettore tridimensionale. La loro implementazione deve essere più flessibile.

Ogni oggetto in Python ha un attributo speciale chiamato `__dict__`, che è un dizionario che archivia gli attributi dell’oggetto.

Rimuovi le chiamate `print` esistenti. Poi, stampa l’attributo `__dict__` dei tuoi vettori `v1` e `v2` per vedere come sono fatti.

# --hints--

Non dovresti avere `print(v2)` e `print(v2.norm())` nel tuo codice.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

Dovresti stampare l’attributo `__dict__` di `v1` e `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
