---
id: 65f8709620b2ce1a62608f5a
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Quando hai bisogno di consultare dinamicamente alcuni attributi partendo da un input di tipo stringa, la funzione integrata `getattr()` è ciò che ti serve. Prende un oggetto come primo argomento e una stringa contenente il nome dell'attributo come secondo argomento.

Inizia a correggere il metodo `__str__` sostituendo la stringa restituita da `__str__()` con un'espressione generatrice che itera sugli attributi dell'oggetto e chiama `getattr()` per ogni attributo `i`.

# --hints--

Dovresti restituire un'espressione generatrice che itera su `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Dovresti restituire un'espressione generatrice che usa `i` come variabile di iterazione.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Dovresti restituire un'espressione generatrice che chiama `getattr(self, i)` per ogni elemento `i` in `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
