---
id: 65f9cb678070ca1668898c70
title: Passaggio 44
challengeType: 20
dashedName: step-44
---

# --description--

Allo stesso modo in cui `__add__` viene chiamato sotto il cofano quando si sommano due oggetti, il metodo `__sub__` viene chiamato implicitamente in caso di sottrazione.

Ora, definisci un metodo `__sub__` vuoto e assegna due parametri: `self` e `other`. All’interno del tuo nuovo metodo, crea un’istruzione if per verificare se `self` e `other` non appartengono alla stessa classe e restituisci `NotImplemented`, come hai fatto in precedenza.

# --hints--

Dovresti definire un metodo `__sub__` all’interno della classe `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

Il tuo metodo `__sub__` dovrebbe avere due parametri, `self` e `other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

Dovresti creare un’istruzione `if` che verifica se `self` e `other` non appartengono alla stessa classe.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

Dovresti restituire `NotImplemented` all’interno del corpo di `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
