---
id: 65f9bc10a9fb1612e066e9e8
title: Passaggio 42
challengeType: 20
dashedName: step-42
---

# --description--

Quando si sommano due vettori, ogni componente di un vettore viene sommata alla stessa componente dell’altro vettore. Per esempio, sommando `(1, 2)` e `(2, 4)` si genera un terzo vettore `(3, 6)`, dove `3` è la somma di `1` e `2` e `6` è la somma di `2` e `4`.

Il dizionario `kwargs` conterrà le coppie chiave-valore necessarie per istanziare un nuovo vettore della stessa classe dei due vettori sommati.

Trasforma il dizionario vuoto in una comprensione di dizionario che itera su `vars(self)` e per ogni chiave (`i`) crea una coppia chiave-valore, dove la chiave è la stessa chiave dell’iterazione corrente e il suo valore è la somma dei valori degli attributi `i` dei due operandi.

# --hints--

Dovresti trasformare il dizionario vuoto assegnato a `kwargs` in una comprensione di dizionario che itera su `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe usare la variabile `i` per iterare su `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe usare `i` come chiave.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe assegnare la somma dell’attributo `i` di `self` e dell’attributo `i` di `other` alla chiave `i` per ogni `i` in `vars(self)`. Usa la funzione `getattr()` per questo.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
