---
id: 65f9b486989cb90ff3e77ac8
title: Passaggio 39
challengeType: 20
dashedName: step-39
---

# --description--

Il parametro `other` del tuo metodo `__add__` rappresenta l'operando posto a destra dell'operatore `+` in un'operazione di addizione.

Vuoi verificare che l'operando a destra non sia un oggetto della stessa classe dell'operando a sinistra (cioè `self`). Crea un'istruzione `if` che lo verifichi controllando il tipo degli operandi.

# --hints--

Dovresti creare un'istruzione `if` che controlli se `type(self)` e `type(other)` sono diversi.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
