---
id: 65f9b17a638f0d0dcce8c354
title: Passaggio 38
challengeType: 20
dashedName: step-38
---

# --description--

Per creare uno spazio vettoriale, devi definire come i vettori dovrebbero comportarsi in diversi casi. I vettori possono essere sommati, formando un nuovo vettore.

Il metodo speciale `__add__` può essere implementato per sovrascrivere ciò che succede per impostazione predefinita quando due oggetti vengono sommati usando l'operatore `+`.

Al momento, provare a sommare due istanze di `R2Vector` o `R3Vector` genererebbe un'eccezione. Crea un metodo `__add__` vuoto all'interno della classe `R2Vector` e dagli due parametri: `self` e `other`.

# --hints--

Dovresti definire un metodo `__add__` all'interno della classe `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

Il tuo metodo `__add__` dovrebbe avere due parametri, `self` e `other`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
