---
id: 65f9dd6e5a08af19c196c2df
title: Passaggio 45
challengeType: 20
dashedName: step-45
---

# --description--

Il vettore risultante dalla sottrazione di un vettore da un altro si ottiene calcolando la differenza di ciascuno dei loro componenti. Per esempio, sottrarre `(2, 4)` da `(7, 3)` genera un terzo vettore `(5, -1)`, dove `5` è la differenza tra `7` e `2` e `-1` è la differenza tra `3` e `4`.

Come hai fatto prima all'interno del metodo `__add__`, dichiara una variabile `kwargs` dopo l'istruzione `if`. Assegnale una comprensione di dizionario che itera sugli attributi dell'oggetto e crea una coppia chiave-valore per ogni chiave `i`, dove la chiave è la stessa chiave dell'iterazione corrente e il suo valore è la differenza tra i valori degli attributi `i` dei due operandi.

# --hints--

Dovresti dichiarare una variabile `kwargs` e assegnarle una comprensione di dizionario che itera sugli attributi dell'oggetto.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe usare la variabile `i` per iterare sugli attributi dell'oggetto.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe usare `i` come chiave.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

La comprensione di dizionario assegnata a `kwargs` dovrebbe assegnare la differenza tra l'attributo `i` di `self` e l'attributo `i` di `other` alla chiave `i` per ogni `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
