---
id: 66682150151af29efec9727d
title: Passaggio 33
challengeType: 20
dashedName: step-33
---

# --description--

Il metodo `__getattribute__` viene chiamato automaticamente ogni volta che provi a consultare un attributo di un'istanza. Se l'attributo non viene trovato a livello di istanza, il metodo lo cercherà a livello di classe e infine nelle classi genitrici.

All'interno della tua classe, definisci un metodo `__getattribute__` con due parametri, `self` e `attr`, e fallo restituire la stringa `'calling __getattribute__'`. Sovrascriverai a breve l'implementazione predefinita solo per vedere come funziona la consultazione degli attributi.

# --hints--

Dovresti definire un metodo chiamato `__getattribute__` con due parametri, `self` e `attr`, all'interno della classe `R2Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

Il tuo metodo `__getattribute__` dovrebbe restituire la stringa `'calling __getattribute__'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
