---
id: 615380dff67172357fcf0425
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

画像が大きすぎます。

`.gallery img` セレクターを作成して画像を指定してください。すべてに `width` を `100%` に、`max-width` を `350px` に設定してください。

また、画像のサイズを均一に保つために `height` プロパティを `300px` に設定してください。

# --hints--

`.gallery img` セレクターが必要です。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery img'));
```

`.gallery img` セレクターには `width` プロパティを `100%` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.width, '100%');
```

`.gallery img` セレクターには `max-width` プロパティを `350px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.maxWidth, '350px');
```

`.gallery img` セレクターには `height` プロパティを `300px` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery img')?.height, '300px');
```

# --seed--


## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```
