---
id: 61537bb9b1a29430ac15ad38
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このプロジェクトは、`styles.css` ファイルをページにリンクすることから始めてください。

# --hints--

`link` 要素の中に `head` 要素があるはずです。

```js
assert.exists(document.querySelector("head>link"));
```

`link` 要素には `rel="stylesheet"` と `href="styles.css"` が必要です。

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
