---
id: 6153908a366afb4d57185c8d
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

Flexboxには主軸と交差軸があります。主軸は`flex-direction`プロパティで定義されており、4つの値が可能です：

- `row`（デフォルト）：左から右へのフレックスアイテムが並ぶ水平軸
- `row-reverse`：右から左へのフレックスアイテムが並ぶ水平軸
- `column`：上から下へのフレックスアイテムが並ぶ垂直軸
- `column-reverse`：下から上へのフレックスアイテムが並ぶ垂直軸

**注意**：軸と方向はテキストの方向によって異なります。ここに示した値は左から右へのテキスト方向の場合です。

異なる値を試して、レイアウトにどのような影響があるか確認してください。

完了したら、`flex-direction`要素に対して明示的に`row`を`.gallery`に設定してください。

# --hints--

`.gallery`セレクターには、`flex-direction`プロパティを`row`の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.flexDirection, 'row');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}


.gallery {
  display: flex;
--fcc-editable-region--
  
--fcc-editable-region--
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
