---
id: 6153893900438b4643590367
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

`.header` のテキストを中央揃えにしてください。`text-transform` プロパティを使い、値に `uppercase` を指定してテキストを大文字にしてください。

すべての辺に `32px` のパディングを設定してください。背景色は `#0a0a23`、テキストの色は `#fff` に設定してください。

`border-bottom` を追加し、値に `4px solid #fdb347` を指定してください。

# --hints--

`.header` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.header'));
```

`.header` セレクターには、`text-align` プロパティを `center` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textAlign, 'center');
```

`.header` セレクターには、`text-transform` プロパティを `uppercase` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.textTransform, 'uppercase');
```

`.header` セレクターには、`padding` プロパティを `32px` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.padding, '32px');
```

`.header` セレクターには、`background-color` プロパティを `#0a0a23` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.backgroundColor, 'rgb(10, 10, 35)');
```

`.header` セレクターには、`color` プロパティを `#fff` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.color, 'rgb(255, 255, 255)');
```

`.header` セレクターには、`border-bottom` プロパティを `4px solid #fdb347` の値で設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.header')?.borderBottom, '4px solid rgb(253, 179, 71)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}
body  {
  margin: 0;
  font-family: sans-serif;
  background-color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
}
```
