---
id: 6494da0daf5df5197963671d
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

青い画像のボーダーが赤いギャラリーのボーダーを超えているのに気づきましたか。これはブラウザーがコンテナ要素のサイズを計算する方法によるものです。

`box-sizing` プロパティはこの動作を設定するために使われます。デフォルトでは `content-box` モデルが使われています。このモデルでは、要素に特定の幅がある場合、その幅は要素のコンテンツのみに基づいて計算されます。パディングとボーダーの値は合計幅に加算されるため、これらの値を収めるために要素のサイズが大きくなります。

グローバルな `box-sizing` セレクターで `content-box` を明示的に `*` に設定してみてください。この時点では、デフォルト値を使っているため、変化は見られません。

# --hints--

`*` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

`*` セレクターには `box-sizing` プロパティがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

`*` セレクターの `box-sizing` プロパティは `content-box` に設定してください。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
