---
id: 6606a219f9efbf38ad496f67
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Sie können auch die Punktnotation verwenden, um auf ein Instanzattribut zuzugreifen.

Außerhalb der `Board`-Klasse, nachdem Sie das `gameboard`-Objekt initialisiert haben, verwenden Sie `gameboard.board`, um auf das `board`-Attribut Ihres `gameboard`-Objekts zuzugreifen und geben Sie das Ergebnis auf dem Bildschirm aus.

# --hints--

Sie sollten `gameboard.board` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.board)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
