---
id: 6606c3fd5634684d48a7887b
title: Schritt 32
challengeType: 20
dashedName: step-32
---

# --description--

Der Generatorausdruck, den Sie im vorherigen Schritt geschrieben haben, erzeugt eine Liste von Boolean-Werten, die darstellen, ob die Bedingung `self.board[row][col] != num` für jedes Element in der angegebenen Spalte über alle Zeilen `True` oder `False` ist.

Übergeben Sie diesen Generatorausdruck an die `all()`-Funktion, um zu überprüfen, ob alle Elemente in der Spalte unterschiedlich von `num` sind, und geben Sie das Ergebnis zurück.

# --hints--

Sie sollten den Generatorausdruck als Argument an einen `all()`-Aufruf übergeben und das Ergebnis zurückgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("return all(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        (self.board[row][col] != num for row in range(9))
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
