---
id: 660699aabc59c532f2d556e5
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Um eine Instanzmethode aufzurufen, müssen Sie die Punktnotation verwenden:

```py
instance_name.method_name()
```

Dabei ist `instance_name` die Instanz oder das Objekt, und `method_name` die Methode, die Sie aufrufen möchten.

Rufen Sie die `spam`-Methode des `gameboard`-Objekts auf.

# --hints--

Sie sollten die `spam`-Methode des `gameboard`-Objekts mit `gameboard.spam()` aufrufen.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
