---
id: 66068fb0bfddba2b7977eb60
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Projekt lernen Sie Klassen und Objekte kennen, indem Sie einen Sudoku-Rätsellöser erstellen.

In Python ist eine Klasse eine Vorlage zur Erstellung von Objekten. Objekte, die aus einer Klasse erstellt werden, sind Instanzen dieser Klasse. Sie können eine Klasse mit folgender Syntax erstellen:

```py
class ClassName:
    pass
```

Dabei ist `class` das Schlüsselwort, das zur Definition der Klasse erforderlich ist, und `ClassName` der Name der Klasse, der konventionell in *PascalCase* geschrieben wird.

Beginnen Sie damit, eine `Board`-Klasse zu erstellen.

# --hints--

Sie sollten eine Klasse mit dem Namen `Board` erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
