---
id: 66069e5759b800364707988e
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Zurück zur `__init__`-Methode: Sie benötigt einen zusätzlichen Parameter, der das zu lösende Puzzle darstellt.

Fügen Sie der `board`-Methode einen zweiten Parameter namens `__init__` hinzu und korrigieren Sie die Instanziierung von `gameboard`, indem Sie ihr die `puzzle`-Liste übergeben, so wie Sie ein Argument an einen Funktionsaufruf übergeben würden.

# --hints--

Ihre `__init__`-Methode sollte zwei Parameter in der Reihenfolge `self` und `board` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self, board")`)) })
```

Ihre `gameboard`-Variable sollte den Wert von `Board(puzzle)` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board(puzzle)")`)) })
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board()
--fcc-editable-region--
```
