---
id: 66069167b3307b2f4067b22b
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Eine neue Instanz einer Klasse wird erstellt, indem die Funktionsnotation verwendet wird, bei der ein Paar runder Klammern an den Klassennamen angehängt wird.

Erstellen Sie außerhalb der Klassendefinition eine Instanz der `Board`-Klasse und weisen Sie sie einer Variablen namens `gameboard` zu.

# --hints--

Sie sollten eine Variable `gameboard` außerhalb der `Board`-Klasse deklarieren.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Ihre `gameboard`-Variable sollte den Wert von `Board()` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
