---
id: 6606c14182435d4bab0de2ee
title: Schritt 30
challengeType: 20
dashedName: step-30
---

# --description--

Als Nächstes erstellen Sie eine Methode, die überprüft, ob eine Zahl in eine bestimmte Spalte des Sudoku-Boards eingefügt werden kann, indem sie prüft, ob die Zahl in dieser Spalte noch nicht vorhanden ist.

Erstellen Sie innerhalb der `Board`-Klasse eine Methode namens `valid_in_col` und geben Sie ihr drei Parameter: `self`, `col` und `num`.
Dabei sind `col` und `num` der Spaltenindex beziehungsweise die zu überprüfende Zahl.

# --hints--

Sie sollten eine neue Methode namens `valid_in_col` innerhalb der `Board`-Klasse erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

Ihre `valid_in_col`-Methode sollte drei Parameter haben: `self`, `col` und `num`, in dieser Reihenfolge.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
