---
id: 66069b0b36053733a2f012fe
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Löschen Sie jetzt Ihren `spam`-Aufruf.

# --hints--

Sie sollten `gameboard.spam()` nicht in Ihrem Code haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
gameboard.spam()
--fcc-editable-region--
```
