---
id: 6606a4641ec48b3a9fe8c2fc
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Die eingebaute Funktion `enumerate` nimmt ein iterierbares Objekt als Argument und gibt ein enumerate-Objekt zurück, das Sie durchlaufen können.
Es liefert die Zählung (die standardmäßig bei null beginnt) und den Wert aus dem iterierbaren Objekt.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

Die Schleife aus dem obigen Beispiel würde die Tuples `0, a`, `1, b` und `2, c` ausgeben.

Ersetzen Sie innerhalb der `find_empty_cell`-Methode `pass` durch eine `for`-Schleife, die die `enumerate()`-Funktion verwendet, um jede Zeile im Sudoku-Brett zu durchlaufen. Verwenden Sie `row` als Index der aktuellen Zeile und `contents` für die Elemente der aktuellen Zeile.


# --hints--

Sie sollten `pass` durch eine `for`-Schleife ersetzen, die `enumerate(self.board)` durchläuft. Verwenden Sie `row` und `contents` als Schleifenvariablen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
