---
id: 66069b992c1c5e3451f3deb0
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Die Instanziierung erzeugt ein leeres Objekt. Die `__init__`-Methode ist eine spezielle Methode, die es Ihnen erlaubt, ein Objekt in einem angepassten Zustand zu instanziieren. Wenn eine Klasse eine `__init__`-Methode implementiert, wird `__init__` automatisch bei der Instanziierung aufgerufen.

Löschen Sie in Ihrer `Board`-Klasse die `spam`-Methode und ersetzen Sie sie durch eine `__init__`-Methode, die einen `self`-Parameter enthält.

# --hints--

Sie sollten keine `spam`-Methode in Ihrer `Board`-Klasse haben.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Sie sollten eine `__init__`-Methode in Ihrer `Board`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Ihre `__init__`-Methode sollte einen `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
