---
id: 6606b224a69a293f98f8db8f
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Die `.index()`-Methode löst eine `ValueError`-Ausnahme aus, wenn der Wert nicht gefunden wird. Um zu verhindern, dass das Programm die Ausführung stoppt, verschachteln Sie diese Codezeile in einem `try`-Block. Die `try`-Anweisung wird verwendet, um Code einzuschließen, der eine Ausnahme auslösen könnte. Die `except`-Klausel hingegen bietet alternativen Code, der ausgeführt wird, falls eine Ausnahme auftritt:

```py
try:
    <code>
except:
    <code>
```

Setzen Sie die Zuweisung von `col` in einen `try`-Block. Erstellen Sie dann eine `except`-Klausel und füllen Sie deren Rumpf mit `pass`.

# --hints--

Sie sollten die Zuweisung von `col` in einen `try`-Block setzen und eine `except`-Klausel erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\nexcept:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            col = contents.index(0)
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
