---
id: 66069f86f58f85371d47123e
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Ein Attribut ist eine Variable, die mit einem Objekt verknüpft ist und verwendet wird, um Daten wie reguläre Variablen zu speichern.

Weisen Sie innerhalb der `__init__`-Methode den `board`-Parameter (der beim Erstellen einer Instanz der `Board`-Klasse übergeben wird) einem Instanzattribut `board` mit `self.board` zu.

`self.board` bezieht sich auf das `board`-Attribut der Instanz der Klasse. Es ist eine Variable, die zum Objekt gehört, das aus der `Board`-Klasse erstellt wurde.

# --hints--

Sie sollten `pass` löschen und den `board`-Parameter innerhalb der `self.board`-Methode `__init__` zuweisen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
