---
id: 6606a3ccb1eea93a23c066bf
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Jetzt arbeiten Sie an einer Methode, die die leeren Felder im Sudoku-Brett findet.

Erstellen Sie innerhalb der `Board`-Klasse eine leere Methode namens `find_empty_cell` und geben Sie ihr einen `self`-Parameter.

# --hints--

Sie sollten eine Methode namens `find_empty_cell` innerhalb Ihrer `Board`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("find_empty_cell")`)) })
```

Ihre `find_empty_cell`-Methode sollte einen Parameter `self` haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        self.board = board
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
