---
id: 6606b961ebcf04460f8af76e
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Testen Sie, ob die `find_empty_cell`-Methode richtig funktioniert, indem Sie sie auf `gameboard` aufrufen und das Ergebnis ausgeben.

Beachten Sie, dass `find_empty_cell` zwar mit einem Parameter definiert ist, Sie ihm jedoch keinen Wert übergeben dürfen, indem Sie ein Argument beim Funktionsaufruf angeben, da `self` automatisch als das Objekt übergeben wird, auf dem Sie die Methode aufrufen.

# --hints--

Sie sollten `gameboard.find_empty_cell()` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.find_empty_cell())")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
