---
id: 6606927d010be4300a4e5330
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Die Instanziierung erzeugt ein leeres Objekt. Aber Klassen können Methoden haben, die wie lokale Funktionen für jede Instanz sind. Innerhalb einer Klasse werden Methoden wie folgt deklariert:

```py
class ClassName:
    def method_name():
        pass
```

Ersetzen Sie innerhalb der `Board`-Klasse `pass` durch eine leere Methode `spam`.

# --hints--

Sie sollten eine Methode namens `spam` innerhalb der `Board`-Klasse definieren.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
