---
id: 6606d8c323d6205890fbbd54
title: Schritt 46
challengeType: 20
dashedName: step-46
---

# --description--

Erstellen Sie innerhalb der `Board`-Klasse eine weitere Methode `is_valid` und geben Sie ihr drei Parameter: `self`, `empty` und `num`. Dabei ist `empty` ein Tuple, das die Zeilen- und Spaltenindizes einer leeren Zelle darstellt, und `num` ist die Zahl, die überprüft werden soll.

Diese Methode prüft, ob eine gegebene Zahl eine gültige Wahl für eine leere Zelle im Sudoku-Brett ist, indem sie die Kompatibilität mit der Zeile, der Spalte und dem 3x3-Quadrat der angegebenen leeren Zelle validiert.

# --hints--

Sie sollten eine neue Methode namens `is_valid` innerhalb der `Board`-Klasse erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

Ihre `is_valid`-Methode sollte drei Parameter haben: `self`, `empty` und `num`, in dieser Reihenfolge.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
