---
id: 660699119472f332798860ad
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Ersetzen Sie jetzt `pass` durch einen `print`-Aufruf und übergeben Sie ihm den String `'Spam!'`.

# --hints--

Sie sollten `pass` löschen und den String `'Spam!'` innerhalb der `spam`-Methode ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
