---
id: 6606b6b7760d0643c3b4eb29
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Wenn der Code innerhalb des `try`-Blocks eine Ausnahme auslöst, soll das Programm weiterlaufen, und die `pass`-Anweisung sorgt genau dafür.

Obwohl dieser Code funktioniert, gilt es als gute Praxis, den Ausnahmetyp nach dem `except`-Schlüsselwort anzugeben.

Da Sie wissen, dass eine `ValueError` ausgelöst werden könnte, lassen Sie nach dem `except`-Schlüsselwort ein Leerzeichen und fügen danach `ValueError` hinzu.

# --hints--

Sie sollten `except ValueError:` in Ihrem Code haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_for_loops()[0].find_bodies()[0].is_equivalent("try:\\n  col = contents.index(0)\\n  return row, col\\nexcept ValueError:\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except:
                pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
