---
id: 6606933d6813a8308c962dd1
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Damit eine Methode eine Instanzmethode ist, benötigt sie einen speziellen Parameter, der gemäß Konvention `self` genannt wird. Dieser Parameter ist eine Referenz auf die Instanz der Klasse und muss immer der erste Parameter sein.

Fügen Sie Ihrer `spam`-Methode einen `self`-Parameter hinzu.

# --hints--

Ihre `spam`-Methode sollte einen `self`-Parameter haben.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
