---
id: 6606c2d203a8124c83b2234b
title: Schritt 31
challengeType: 20
dashedName: step-31
---

# --description--

Sie müssen überprüfen, ob eine gegebene Zahl ungleich der Zahl in der angegebenen Spalte der aktuellen Zeile ist.

Ersetzen Sie dazu `pass` durch einen Generatorausdruck, der über den Bereich von `0` bis `8` (einschließlich) iteriert und für jedes `row` auswertet, ob die Zahl an der angegebenen `row` und Spalte `col` auf dem Spielfeld ungleich `num` ist.

# --hints--

Sie sollten `pass` löschen und einen Generatorausdruck `(self.board[row][col] != num for row in range(9))` erstellen.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").find_body().is_equivalent("(self.board[row][col] != num for row in range(9))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    def valid_in_col(self, col, num):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
