---
id: 6606bf4561f8794a0d345919
title: Schritt 27
challengeType: 20
dashedName: step-27
---

# --description--

Es ist Zeit, die `valid_in_row`-Methode zu testen. Rufen Sie `valid_in_row` auf `gameboard` auf. Übergeben Sie ihr `0` und `8` als Argumente und geben Sie das Ergebnis aus.

Beachten Sie erneut, wie die Methode mit drei Parametern definiert ist, aber nur mit zwei Argumenten aufgerufen wird, da `self` automatisch als das Objekt übergeben wird, auf dem die Methode aufgerufen wird.

# --hints--

Sie sollten das Ergebnis des Aufrufs von `valid_in_row(0, 8)` auf `gameboard` ausgeben.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(gameboard.valid_in_row(0, 8))")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--

--fcc-editable-region--
```
