---
id: 660a737f0f72b51de361051c
title: Schritt 54
challengeType: 20
dashedName: step-54
---

# --description--

Der `:=`-Operator ermöglicht es Ihnen, Variablen mitten in einem Ausdruck zuzuweisen. Die Syntax lautet: `name := val`, wobei `name` der Variablenname und `val` der Variablenwert ist.

Diese Konstruktion wird formal als *Zuweisungsausdruck* bezeichnet, während der `:=`-Operator gemeinhin als *Walross*-Operator bekannt ist.

Da Sie den `self.find_empty_cell()`-Aufruf mehr als einmal benötigen, weisen Sie ihn mit dem Walross-Operator einer Variablen `next_empty` zu. Schließen Sie dann die Zuweisung in ein Paar runder Klammern ein.

Auf diese Weise kombinieren Sie die Zuweisung und die Bedingungsprüfung in einer einzigen Zeile, wodurch der Code kompakter wird.

# --hints--

Sie sollten die `if`-Bedingung in `(next_empty := self.find_empty_cell()) is None` ändern.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("solver").find_ifs()[0].find_conditions()[0].is_equivalent("(next_empty := self.find_empty_cell()) is None")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True

    def is_valid(self, empty, num):
        row, col = empty
        valid_in_row = self.valid_in_row(row, num)
        valid_in_col = self.valid_in_col(col, num)
        valid_in_square = self.valid_in_square(row, col, num)
        return all([valid_in_row, valid_in_col, valid_in_square])
--fcc-editable-region--
    def solver(self):
        if self.find_empty_cell() is None:
            return True
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
