---
id: 66fbcf750a62784cf11f562d
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`recipe1`라는 이름의 객체를 만드세요. `recipe1` 객체 안에 `name` 속성을 만들고 값으로 `"Spaghetti Carbonara"`을 넣으세요.

또한 `recipe1` 객체 안에 배열을 값으로 하는 `ingredients` 속성을 만드세요. 배열 안에는 `spaghetti`, `Parmesan cheese`, `pancetta`, 그리고 `black pepper`가 들어 있어야 합니다.

# --hints--

`recipe1`라는 이름의 객체를 만들어야 합니다.

```js
assert.isObject(recipe1);
```

`recipe1` 객체는 `name` 속성을 가져야 합니다.

```js
assert.property(recipe1, 'name');
```

`name` 속성은 `Spaghetti Carbonara`로 설정해야 합니다.

```js
assert.equal(recipe1.name, 'Spaghetti Carbonara');
```

`recipe1` 객체는 `ingredients` 속성을 가져야 합니다.

```js
assert.property(recipe1, 'ingredients');
```

`ingredients` 속성은 배열 값을 가져야 합니다.

```js
assert.isArray(recipe1.ingredients);
```

`ingredients` 속성의 배열 값에는 `spaghetti`, `Parmesan cheese`, `pancetta`, 그리고 `black pepper`가 포함되어야 합니다.

```js
assert.deepEqual(recipe1.ingredients, ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"]);
```

# --seed--

## --seed-contents--

```js
const recipes = [];
--fcc-editable-region--

--fcc-editable-region--
```
