---
id: 66fbcf750a62784cf11f5630
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 세 개의 객체를 `recipes` 배열에 추가해야 합니다. 이를 위해 `push()` 메서드를 사용할 수 있습니다.

`push()` 메서드를 사용해 모든 레시피 객체를 `recipes` 배열에 추가하세요. `recipe1`, `recipe2`, `recipe3` 순서대로 추가해야 합니다.

또한 `recipe1Name`, `recipe2CookingTime`, `recipe3Ingredients` 변수를 삭제하고, 해당 변수를 출력하는 `console.log` 문장도 삭제하세요.

# --hints--

모든 레시피 객체를 `recipes` 배열에 추가해야 합니다. 선언된 순서대로 추가하세요.

```js
assert.deepEqual(recipes, [recipe1, recipe2, recipe3]);
```

`recipe1Name` 변수를 삭제하세요.

```js
assert.throws(() => recipe1Name, ReferenceError);
```

`console.log(recipe1Name);` 문장을 삭제하세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe1Name\s*\)\s*;?/);
```

`recipe2CookingTime` 변수를 삭제하세요.

```js
assert.throws(() => recipe2CookingTime, ReferenceError);
```

`console.log(recipe2CookingTime);` 문장을 삭제하세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe2CookingTime\s*\)\s*;?/);
```

`recipe3Ingredients` 변수를 삭제하세요.

```js
assert.throws(() => recipe3Ingredients, ReferenceError);
```

`console.log(recipe3Ingredients);` 문장을 삭제하세요.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*recipe3Ingredients\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--
const recipe1Name = recipe1.name;
console.log(recipe1Name);

const recipe2CookingTime = recipe2.cookingTime;
console.log(recipe2CookingTime);

const recipe3Ingredients = recipe3.ingredients;
console.log(recipe3Ingredients);
--fcc-editable-region--
```
