---
id: 66fbcdfed6848e48947dbbc9
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript 객체를 사용해 레시피 추적기를 만드세요.

먼저 `recipes`라는 빈 배열을 만드세요. 나중에 레시피 객체를 이 배열에 추가할 것입니다.

# --hints--

`recipes`라는 배열을 만들어야 합니다.

```js
assert.isArray(recipes);
```

`recipes` 배열은 비어 있어야 합니다.

```js
assert.isEmpty(recipes);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
