---
id: 66fbcf750a62784cf11f5636
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

`totalIngredients`와 `difficultyLevel`의 `recipe2`와 `recipe3` 속성에 대해 같은 과정을 반복하세요.

# --hints--

`totalIngredients`의 `recipe2` 속성에 접근해야 합니다.

```js
assert.isNotNull(recipe2.totalIngredients);
```

`getTotalIngredients`의 `recipe2.ingredients` 속성에 `totalIngredients`를 인자로 `recipe2`를 호출한 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe2.totalIngredients, getTotalIngredients(recipe2.ingredients));
```

`difficultyLevel`의 `recipe2` 속성에 접근해야 합니다.

```js
assert.isNotEmpty(recipe2.difficultyLevel);
```

`getDifficultyLevel`의 `recipe2.cookingTime` 속성에 `difficultyLevel`를 인자로 `recipe2`를 호출한 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe2.difficultyLevel, getDifficultyLevel(recipe2.cookingTime));
```

`totalIngredients`의 `recipe3` 속성에 접근해야 합니다.

```js
assert.isNotNull(recipe3.totalIngredients);
```

`getTotalIngredients`의 `recipe3.ingredients` 속성에 `totalIngredients`를 인자로 `recipe3`를 호출한 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe3.totalIngredients, getTotalIngredients(recipe3.ingredients));
```

`difficultyLevel`의 `recipe3` 속성에 접근해야 합니다.

```js
assert.isNotEmpty(recipe3.difficultyLevel);
```

`getDifficultyLevel`의 `recipe3.cookingTime` 속성에 `difficultyLevel`를 인자로 `recipe3`를 호출한 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe3.difficultyLevel, getDifficultyLevel(recipe3.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

recipe1.totalIngredients = getTotalIngredients(recipe1.ingredients);
recipe1.difficultyLevel = getDifficultyLevel(recipe1.cookingTime);

--fcc-editable-region--

--fcc-editable-region--
```
