---
id: 66fbcf750a62784cf11f562e
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

다음 속성과 값을 가진 `recipe2` 객체를 만드세요:

| 키 | 값   |
| ----------- | ------- |
| `name`       | `Chicken Curry` |
| `ingredients`       | `["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]` |
| `cookingTime`       | `42` |
| `totalIngredients`   | `null`   |
| `difficultyLevel`   | `""`   |

# --hints--

`recipe2`라는 이름의 객체를 만드세요.

```js
assert.isObject(recipe2);
```

`recipe2` 객체는 `name` 속성을 가져야 합니다.

```js
assert.property(recipe2, "name");
```

`name` 속성은 `Chicken Curry`로 설정해야 합니다.

```js
assert.equal(recipe2.name, "Chicken Curry");
```

`recipe2` 객체는 `ingredients` 속성을 가져야 합니다.

```js
assert.property(recipe2, "ingredients");
```

`ingredients` 속성은 배열 값을 가져야 합니다.

```js
assert.isArray(recipe2.ingredients);
```

`ingredients` 속성의 배열 값에는 `chicken breast`, `coconut milk`, `curry powder`, `onion` 그리고 `garlic`가 포함되어야 합니다.

```js
assert.deepEqual(recipe2.ingredients, ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"]);
```

`recipe2` 객체는 `cookingTime` 속성을 가져야 합니다.

```js
assert.property(recipe2, "cookingTime");
```

`cookingTime` 속성 값은 숫자여야 합니다.

```js
assert.isNumber(recipe2.cookingTime);
```

`cookingTime` 속성은 `42`로 설정해야 합니다.

```js
assert.equal(recipe2.cookingTime, 42);
```

`recipe2` 객체는 `totalIngredients` 속성을 가져야 합니다.

```js
assert.property(recipe2, "totalIngredients");
```

`totalIngredients` 속성은 `null`로 설정해야 합니다.

```js
assert.isNull(recipe2.totalIngredients);
```

`recipe2` 객체는 `difficultyLevel` 속성을 가져야 합니다.

```js
assert.property(recipe2, "difficultyLevel");
```

`difficultyLevel` 속성은 빈 문자열로 설정해야 합니다.

```js
assert.deepEqual(recipe2.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
  };

--fcc-editable-region--

--fcc-editable-region--
```
