---
id: 66fbcf750a62784cf11f5632
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

재료가 들어 있는 배열을 나타내는 단일 인수를 받고, 함수에 전달된 배열에서 재료 개수를 반환하는 `getTotalIngredients` 함수를 만드세요.

# --hints--

`getTotalIngredients` 함수를 만드세요.

```js
assert.isFunction(getTotalIngredients);
```

`getTotalIngredients` 함수는 단일 매개변수를 가져야 합니다.

```js
assert.lengthOf(getTotalIngredients, 1);
```

`getTotalIngredients` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getTotalIngredients(recipe1.ingredients));
```

`getTotalIngredients(recipe1.ingredients)`는 `4`를 반환해야 합니다.

```js
assert.strictEqual(getTotalIngredients(recipe1.ingredients), 4);
```

`getTotalIngredients(recipe2.ingredients)`는 `5`를 반환해야 합니다.

```js
assert.strictEqual(getTotalIngredients(recipe2.ingredients), 5);
```

`getTotalIngredients(recipe3.ingredients)`는 `3`를 반환해야 합니다.

```js
assert.strictEqual(getTotalIngredients(recipe3.ingredients), 3);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

--fcc-editable-region--

--fcc-editable-region--
```
