---
id: 66fbcf750a62784cf11f5634
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

각 함수를 테스트할 시간입니다. 어떤 레시피를 사용해도 되지만, 이 튜토리얼은 `recipe1`부터 시작합니다.

두 개의 새 변수를 만드세요: `recipe1TotalIngredients`과 `recipe1DifficultyLevel`. 각 변수에 해당하는 함수를 호출하고 적절한 `recipe1` 속성을 전달하여 값을 할당하세요.

마지막으로, 결과를 확인하기 위해 각 변수를 콘솔에 출력하세요.

# --hints--

`recipe1TotalIngredients` 변수를 만들어야 합니다.

```js
assert.isNotNull(recipe1TotalIngredients);
```

`recipe1TotalIngredients` 변수는 `getTotalIngredients(recipe1.ingredients)`로 설정해야 합니다.

```js
assert.equal(recipe1TotalIngredients, getTotalIngredients(recipe1.ingredients));
```

`recipe1TotalIngredients`을 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(recipe1TotalIngredients\);?/);
```

`recipe1DifficultyLevel` 변수를 만들어야 합니다.

```js
assert.isNotNull(recipe1DifficultyLevel);
```

`recipe1DifficultyLevel` 변수는 `getDifficultyLevel(recipe1.cookingTime)`로 설정해야 합니다.

```js
assert.equal(recipe1DifficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

`recipe1DifficultyLevel`을 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(recipe1DifficultyLevel\);?/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
