---
id: 66fbcf750a62784cf11f5633
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

요리 시간을 나타내는 숫자를 매개변수로 받는 `getDifficultyLevel` 함수를 만드세요.

요리 시간이 `30` 이하이면 함수는 `"easy"`를 반환해야 합니다. 요리 시간이 `60` 이하이면 함수는 `"medium"`를 반환해야 합니다. 그렇지 않으면 함수는 `"hard"`를 반환해야 합니다.

# --hints--

`getDifficultyLevel` 함수를 만드세요.

```js
assert.isFunction(getDifficultyLevel);
```

`getDifficultyLevel` 함수는 매개변수를 하나 가져야 합니다.

```js
assert.lengthOf(getDifficultyLevel, 1);
```

요리 시간이 `getDifficultyLevel` 이하일 때 `"easy"` 함수는 `30`를 반환해야 합니다.

```js
assert.strictEqual(getDifficultyLevel(10), "easy");
assert.strictEqual(getDifficultyLevel(20), "easy");
assert.strictEqual(getDifficultyLevel(29), "easy");
assert.strictEqual(getDifficultyLevel(30), "easy");
```

요리 시간이 `getDifficultyLevel`보다 크고 `"medium"` 이하일 때 `31` 함수는 `60`를 반환해야 합니다.

```js
assert.strictEqual(getDifficultyLevel(31), "medium");
assert.strictEqual(getDifficultyLevel(40), "medium");
assert.strictEqual(getDifficultyLevel(50), "medium");
assert.strictEqual(getDifficultyLevel(60), "medium");
```

요리 시간이 `getDifficultyLevel`보다 클 때 `"hard"` 함수는 `60`를 반환해야 합니다.

```js
assert.strictEqual(getDifficultyLevel(61), "hard");
assert.strictEqual(getDifficultyLevel(65), "hard");
assert.strictEqual(getDifficultyLevel(70), "hard");
assert.strictEqual(getDifficultyLevel(75), "hard");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

--fcc-editable-region--

--fcc-editable-region--
```
