---
id: 66fbcf750a62784cf11f5635
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 `recipes` 배열의 각 항목에 `totalIngredients`과 `difficultyLevel` 속성에 대한 값을 채울 수 있습니다.

우선 `totalIngredients`의 `difficultyLevel`과 `recipe1`에 접근하여 함수 호출과 인수의 적절한 결과로 설정하세요.

# --hints--

`totalIngredients`의 `recipe1` 속성에 접근해야 합니다.

```js
assert.isNotNull(recipe1.totalIngredients);
```

`getTotalIngredients`의 `recipe1.ingredients` 속성에 `totalIngredients`를 인수로 호출한 `recipe1`의 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe1.totalIngredients, getTotalIngredients(recipe1.ingredients));
```

`difficultyLevel`의 `recipe1` 속성에 접근해야 합니다.

```js
assert.isNotEmpty(recipe1.difficultyLevel);
```

`getDifficultyLevel`의 `recipe1.cookingTime` 속성에 `cookingTime`를 인수로 호출한 `recipe1`의 결과를 할당해야 합니다.

```js
assert.strictEqual(recipe1.difficultyLevel, getDifficultyLevel(recipe1.cookingTime));
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

recipes.push(recipe1, recipe2, recipe3);

function getTotalIngredients(ingredients) {
  return ingredients.length;
}

function getDifficultyLevel(cookingTime) {
  if (cookingTime <= 30) {
    return "easy";
  } else if (cookingTime <= 60) {
    return "medium";
  } else {
    return "hard";
  }
}

const recipe1TotalIngredients = getTotalIngredients(recipe1.ingredients);
console.log(recipe1TotalIngredients);

const recipe1DifficultyLevel = getDifficultyLevel(recipe1.cookingTime);
console.log(recipe1DifficultyLevel);

--fcc-editable-region--

--fcc-editable-region--
```
