---
id: 66ff02792b3042a2b811f475
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

다음 속성들을 `recipe1` 객체에 추가하세요:

| 키 | 값   |
| ----------- | ------- |
|`cookingTime`|`22`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

`null` 속성과 빈 문자열 값은 나중에 계산한 후 업데이트할 예정입니다.

# --hints--

`recipe1` 객체에 `cookingTime` 속성이 있어야 합니다.

```js
assert.property(recipe1, "cookingTime");
```

`cookingTime` 속성은 `22`로 설정해야 합니다.

```js
assert.equal(recipe1.cookingTime, 22);
```

`recipe1` 객체에 `totalIngredients` 속성이 있어야 합니다.

```js
assert.property(recipe1, "totalIngredients");
```

`totalIngredients` 속성은 `null`로 설정해야 합니다.

```js
assert.isNull(recipe1.totalIngredients);
```

`recipe1` 객체에 `difficultyLevel` 속성이 있어야 합니다.

```js
assert.property(recipe1, "difficultyLevel");
```

`difficultyLevel` 속성은 빈 문자열로 설정해야 합니다.

```js
assert.deepEqual(recipe1.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
--fcc-editable-region--
  
--fcc-editable-region--
};
```
