---
id: 670e4f45f7116c0f216a5177
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

다음 단계로 넘어가기 전에 객체에서 속성에 접근하는 방법을 연습하세요.

점 표기법(`.`)이나 대괄호 표기법(`[]`) 중 어느 쪽을 사용해도 됩니다. 예시는 다음과 같습니다:

```js
const person = {
  name: "John",
  age: 30,
  job: "Software Engineer"
};

console.log(person.name); // John
console.log(person['age']);  // 30
```

`name`의 `recipe1` 속성에 접근하여 변수 `recipe1Name`에 할당하세요.

다음으로 `cookingTime`의 `recipe2` 속성에 접근하여 변수 `recipe2CookingTime`에 할당하세요.

마지막으로 `ingredients`의 `recipe3` 속성에 접근하여 변수 `recipe3Ingredients`에 할당하세요.

생성한 모든 변수를 콘솔에 출력하는 것을 잊지 마세요.

# --hints--

`recipe1Name` 변수를 생성하세요.

```js
assert.isDefined(recipe1Name);
```

`name`의 `recipe1` 속성 값을 `recipe1name` 변수에 할당하세요.

```js
assert.strictEqual(recipe1Name, recipe1.name);
```

`recipe1Name`를 콘솔에 출력하세요.

```js
assert.match(code, /console\.log\(\s*recipe1Name\s*\)/);
```

`recipe2CookingTime` 변수를 생성하세요.

```js
assert.isDefined(recipe2CookingTime);
```

`cookingTime`의 `recipe2` 속성 값을 `recipe2CookingTime` 변수에 할당하세요.

```js
assert.strictEqual(recipe2CookingTime, recipe2.cookingTime);
```

`recipe2CookingTime`를 콘솔에 출력하세요.

```js
assert.match(code, /console\.log\(recipe2CookingTime\)/);
```

`recipe3Ingredients` 변수를 생성하세요.

```js
assert.isDefined(recipe3Ingredients);
```

`ingredients`의 `recipe3` 속성 값을 `recipe3Ingredients` 변수에 할당하세요.

```js
assert.strictEqual(recipe3Ingredients, recipe3.ingredients);
```

`recipe3Ingredients`를 콘솔에 출력하세요.

```js
assert.match(code, /console\.log\(recipe3Ingredients\)/);
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe3 = {
  name: "Vegetable Stir Fry",
  ingredients: ["broccoli", "carrot", "bell pepper"],
  cookingTime: 15,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
