---
id: 69161f4c0c9c8c48b3c843c5
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

다음 속성과 값을 가진 `recipe3` 객체를 만드세요:

|키|값|
|-----------|-------|
|`name`|`Vegetable Stir Fry`|
|`ingredients`|`["broccoli", "carrot", "bell pepper"]`|
|`cookingTime`|`15`|
|`totalIngredients`|`null`|
|`difficultyLevel`|`""`|

# --hints--

`recipe3`라는 이름의 객체를 만드세요.

```js
assert.isObject(recipe3);
```

`recipe3` 객체는 `name` 속성을 가져야 합니다.

```js
assert.property(recipe3, "name");
```

`name` 속성은 `Vegetable Stir Fry`로 설정해야 합니다.

```js
assert.equal(recipe3.name, "Vegetable Stir Fry");
```

`recipe3` 객체는 `ingredients` 속성을 가져야 합니다.

```js
assert.property(recipe3, "ingredients");
```

`ingredients` 속성은 배열 값을 가져야 합니다.

```js
assert.isArray(recipe3.ingredients);
```

`ingredients` 속성의 배열 값에는 `broccoli`, `carrot`, `bell pepper`가 포함되어야 합니다.

```js
assert.deepEqual(recipe3.ingredients, ["broccoli", "carrot", "bell pepper"]);
```

`recipe3` 객체는 `cookingTime` 속성을 가져야 합니다.

```js
assert.property(recipe3, "cookingTime");
```

`cookingTime` 속성 값은 숫자여야 합니다.

```js
assert.isNumber(recipe3.cookingTime);
```

`cookingTime` 속성은 `15`로 설정해야 합니다.

```js
assert.equal(recipe3.cookingTime, 15);
```

`recipe3` 객체는 `totalIngredients` 속성을 가져야 합니다.

```js
assert.property(recipe3, "totalIngredients");
```

`totalIngredients` 속성은 `null`로 설정해야 합니다.

```js
assert.isNull(recipe3.totalIngredients);
```

`recipe3` 객체는 `difficultyLevel` 속성을 가져야 합니다.

```js
assert.property(recipe3, "difficultyLevel");
```

`difficultyLevel` 속성은 빈 문자열로 설정해야 합니다.

```js
assert.deepEqual(recipe3.difficultyLevel, "");
```

# --seed--

## --seed-contents--

```js
const recipes = [];

const recipe1 = {
  name: "Spaghetti Carbonara",
  ingredients: ["spaghetti", "Parmesan cheese", "pancetta", "black pepper"],
  cookingTime: 22,
  totalIngredients: null,
  difficultyLevel: ""
};

const recipe2 = {
  name: "Chicken Curry",
  ingredients: ["chicken breast", "coconut milk", "curry powder", "onion", "garlic"],
  cookingTime: 42,
  totalIngredients: null,
  difficultyLevel: ""
};

--fcc-editable-region--

--fcc-editable-region--
```
