---
id: 61437d575fb98f57fa8f7f36
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with your standard HTML boilerplate. Add a `DOCTYPE` declaration, an `html` element specifying this page is in English, a `head` element, and a `body` element.

قم باضافة وسم `<meta>` مع الـ `charset` المناسبة و وسم `<meta>` للتجاوب مع شاشات المحمول داخل عنصر `head`.

# --hints--

يجب أن يحتوي الكود الخاص بك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

عليك أن تضيف مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

عليك أن تعرّف أن نوع مستندك هو `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

ينبغي إغلاق إعلان `DOCTYPE` بواسطة <code dir="ltr">&gt;</code> بعد تعيين النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يكون لديك علامة فتح `<html>`، ويجب أن يحتوي على `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصرك `html` على علامة أغلاق (closing tag).

```js
assert(code.match(/<\/html\s*>/));
```

تعريف `DOCTYPE` الخاص بك يجب أن يكون عند بداية HTML الخاص بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك الـ opening tag الآتي `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك الـ closing tag الآتي `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك الـ opening tag الآتي `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك الـ closing tag الآتي `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عناصر `head` و `body` متلازمين (siblings).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

عنصر `head` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

عنصر `body` يجب أن يكون داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

يجب أن يكون لديك عنصران `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

احد عناصر `meta` يجب أن يحتوي على `name` تم تعيينه إلى `viewport`، و `content` تم تعيينه إلى `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

عنصر `meta` الآخر يجب أن يحتوي على السمة `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
