---
id: 614385513d91ae5c251c2052
title: الخطوة الثانية
challengeType: 0
dashedName: step-2
---

# --description--

اضف عنصر `title` يحتوي على النص `Magazine`، وعنصر `link` يحمل خط `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` الى ‏‏ورقة الأنماط (stylesheet), ثم `link` يحمل أدوات `https://use.fontawesome.com/releases/v5.8.2/css/all.css` الى ‏‏ورقة الأنماط (stylesheet), وأستخدم `link` يحمل `./styles.css` لربط ورقة الأنماط (stylesheet).

ستحمل ورقة نمط الخط (stylesheet) الخاصة بك ثلاثة خطوط منفصلة: `Anton`, و `Baskervville`, و `Raleway`.

# --hints--

Your code should have three `link` elements.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

Your `link` elements should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

العناصر `link` الثلاث الخاص بك يجب أن تحتوي على سمة (attribute) الـ `rel` بقيمة `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

أحد عناصر الرابط (link) الخاص بك يجب أن يحتوى على `href` محدد إلى `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

واحد من الروابط (link) الخاص بك يجب أن يحتوي على `href` بقيمة `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

One of your `link` elements should have the `href` set to `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

الكود الخاص بك يجب أن يحتوي على عنصر `title`.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

Your code should have one `title` element.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

Your `title` element should be inside your `head` element.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

يجب أن يكون لمشروعك العنوان `Magazine`. تذكر أن الـحالة الحرف (casing) سواء كبيرا أو صغيرا (capital or small), والإملاء (spelling) مهمان للعنوان(title).

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
