---
id: a26cbbe9ad8655a977e1ceb5
title: створіть застосунок для пошуку найдовшого слова
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

У цій лабораторній роботі ви створите функцію, яка повертає довжину найдовшого слова у заданому реченні.

Наприклад, у реченні `"The quick brown fox jumped over the lazy dog"` найдовшим словом є `"jumped"`, довжина якого дорівнює 6.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні створити функцію з іменем `findLongestWordLength`, яка приймає рядок як параметр.
2. Функція має повертати довжину найдовшого слова у рядку.

# --hints--

Ви повинні створити функцію з іменем `findLongestWordLength`.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` має мати один параметр.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` має повертати число.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` має повертати `6`.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` має повертати `5`.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` має повертати `6`.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` має повертати `8`.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` має повертати `8`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` має повертати `19`.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
