---
id: 5cc0bd7a49b71cb96132e54c
title: بناء مولد نطاق أعداد
challengeType: 26
dashedName: build-a-range-of-numbers-generator
---

# --description--

في هذا المختبر، ستبني دالة `rangeOfNumbers` التي تُولد مصفوفة من الأعداد ضمن نطاق محدد.

أمثلة:

- تُعيد `rangeOfNumbers(3, 9)` المصفوفة `[3, 4, 5, 6, 7, 8, 9]`
- تُعيد `rangeOfNumbers(5, 5)` المصفوفة `[5]`

المتطلبات:

- استخدم العودية (يجب أن تستدعي الدالة نفسها)
- لا تستخدم الحلقات أو طرق المصفوفات (`for`، `while`، `Array.from`، `forEach`، `map`، `filter`، `reduce`)

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `rangeOfNumbers` تأخذ معلمتين: `startNum` و `endNum`.
1. يجب أن تُعيد الدالة مصفوفة من أعداد صحيحة تبدأ بالعدد المُمثل بالمعلمة `startNum` وتنتهي بالعدد المُمثل بالمعلمة `endNum` (شاملة).
1. سيكون `startNum` دائمًا أقل من أو يساوي `endNum`.
1. يجب أن تستخدم دالتك العودية باستدعاء نفسها. لا يجب أن تستخدم أي صيغة حلقات (`for`، `while`، `Array.from()`، أو دوال رتبة عليا مثل `forEach`، `map`، `filter`، أو `reduce`).
1. يجب أن تتعامل الدالة مع الحالة الأساسية عندما يكون `startNum` مساويًا لـ `endNum` بإرجاع مصفوفة تحتوي ذلك العدد فقط.
1. في الحالة العودية، يجب أن تستدعي الدالة نفسها بمعلمات معدلة لبناء المصفوفة، ثم تضيف العدد الحالي إلى النتيجة.
1. يجب ألا تعتمد الدالة على متغيرات عامة لتخزين أو بناء المصفوفة.

# --hints--

يجب أن يكون لديك دالة باسم `rangeOfNumbers`.

```js
assert.isFunction(rangeOfNumbers);
```

يجب أن تأخذ دالة `rangeOfNumbers` معلمتين.

```js
assert.match(rangeOfNumbers.toString(), /rangeOfNumbers\s*\(\s*\w+\s*,\s*\w+\s*\)/);
```

يجب أن تُعيد دالتك مصفوفة.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

يجب ألا يستخدم شفرتك أي صيغة حلقات (`for`، `while`، `Array.from()` أو دوال رتبة عليا مثل `forEach`، `map`، `filter`، أو `reduce`).

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce|Array\.from/g)
);
```

يجب أن تستخدم `rangeOfNumbers` العودية (استدعاء نفسها) لحل هذا التحدٍّ.

```js
assert.isAtLeast(
  __helpers.removeJSComments(rangeOfNumbers.toString()).match(/rangeOfNumbers\s*\(.+\)/g)?.length,
  2
);
```

يجب أن تُعيد `rangeOfNumbers(1, 5)` المصفوفة `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

يجب أن تُعيد `rangeOfNumbers(6, 9)` المصفوفة `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

يجب أن تُعيد `rangeOfNumbers(4, 4)` المصفوفة `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

يجب أن تُعيد `rangeOfNumbers(10, 15)` المصفوفة `[10, 11, 12, 13, 14, 15]`.

```js
assert.deepStrictEqual(rangeOfNumbers(10, 15), [10, 11, 12, 13, 14, 15]);
```

يجب أن تُعيد `rangeOfNumbers(2, 8)` المصفوفة `[2, 3, 4, 5, 6, 7, 8]`.

```js
assert.deepStrictEqual(rangeOfNumbers(2, 8), [2, 3, 4, 5, 6, 7, 8]);
```

يجب ألا تُستخدم المتغيرات العامة لتخزين المصفوفة.

```js
rangeOfNumbers(1, 3);
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
