---
id: a97fd23d9b809dac9921074f
title: Construye una función de suma con argumentos opcionales
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

En este laboratorio construirás una función que acepta hasta dos argumentos y los suma, pero si solo hay un argumento, devuelve una función que espera el segundo número para sumar.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Crea una función `addTogether`.
2. Si la función recibe dos argumentos, la función debe devolver la suma de los dos argumentos.
3. Si solo se proporciona un argumento, `addTogether` debe devolver una función.
   - Cuando la función retornada se llama con un solo argumento, debe devolver la suma.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Si alguno de los argumentos no es un número válido, devuelve `undefined`.

# --hints--

Debes tener la función `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` debe devolver 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` debe devolver 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` debe devolver `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` Debería devolver `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` debe devolver `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` debe devolver una función.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` debe devolver 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` debe devolver `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` debe devolver `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
