---
id: 66def5467aee701733aaf8cc
title: Construir una aplicación de Administrador de Marcadores
challengeType: 25
dashedName: build-a-bookmark-manager-app
demoType: onClick
---

# --description--

Para este laboratorio, todo el HTML y CSS te han sido proporcionados. Usarás JavaScript para completar la aplicación del Administrador de Marcadores para que pueda almacenar una colección de marcadores en el almacenamiento local y leer datos de él.

**Objetivo:** Cumplir con los siguientes requisitos del usuario y lograr que pasen todas las pruebas para completar el laboratorio.

**Requisitos del usuario:**

1. Debe tener una función `getBookmarks` que devuelva el array `bookmarks` almacenado en el almacenamiento local.  Si la clave `bookmarks` aún no se ha configurado o no contiene un array válido con marcadores, la función `getBookmarks` debe devolver un array vacío.
1. La clave `bookmarks` almacenada en el almacenamiento local debe ser un array de objetos de marcador.  Cada objeto de marcador debe tener tres claves: `name`, `category` y `url`.
1. Debe tener una función llamada `displayOrCloseForm` que alterna la clase `hidden` entre `#main-section` y `#form-section`.
1. Al hacer clic en `#add-bookmark-button`, debe actualizar el texto interno de `.category-name` al valor de la opción seleccionada de `#category-dropdown` y llamar a `displayOrCloseForm` para mostrar la sección del formulario y ocultar la sección principal.
1. Cuando haga clic en `#close-form-button`, debe ejecutar su función para ocultar la sección del formulario y mostrar la sección principal.
1. Al hacer clic en `#add-bookmark-button-form`, debe actualizar la clave `bookmarks` almacenada en el almacenamiento local agregando un objeto de marcador al final del array. El objeto debe tener `name` establecido en el valor del input `#name`, `category` establecido en el valor de la opción seleccionada del desplegable de categorías, y `url` establecido en el valor del input `#url`.
1. Una vez que se actualice la clave `bookmarks`, debe restablecer el valor de `#name` y `#url` a una cadena vacía antes de ejecutar su función para ocultar la sección del formulario y mostrar la sección principal.
1. Debe tener una función llamada `displayOrHideCategory` que alterna la clase `hidden` en `#main-section` y `#bookmark-list-section`.
1. Cuando haga clic en `#view-category-button`, debe actualizar el texto interno de `.category-name` al valor de la opción seleccionada de `#category-dropdown`, modificar el HTML interno de `#category-list` según las historias de usuario a continuación, y llamar a la función `displayOrHideCategory`.
1. Si ninguno de los marcadores en el almacenamiento local tiene la categoría, debe establecer el HTML interno de `#category-list` en un elemento `p` con el texto `No se encontraron marcadores`.
1. Si uno o más marcadores en el almacenamiento local tienen la categoría seleccionada, agregue un botón de radio para cada marcador al `innerHTML` de `#category-list`. Dé a cada botón de radio un `id` y `value` establecidos en el nombre del marcador, y un `name` que sea el mismo para todos los botones.
1. Cada botón de radio debe tener una etiqueta correspondiente que contenga un elemento de anclaje con el nombre del marcador y el atributo `href` establecido en la URL del marcador.
1. Al hacer clic en el `#close-list-button`, debe ejecutar su función para ocultar el `#bookmark-list-section` y mostrar la sección principal.
1. Al hacer clic en el `#delete-bookmark-button`, debe eliminar el marcador correspondiente al botón de radio seleccionado y la categoría adecuada del almacenamiento local y actualizar la lista de marcadores mostrada.

# --antes de todo--

```js
const getHidden = (notMain) => {
  const mainSection = document.querySelector("#main-section");
  const formSection = document.querySelector("#form-section");
  const bookmarkListSection = document.querySelector("#bookmark-list-section");
  if (notMain === "form") {
    if (mainSection.classList.contains("hidden") && !formSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && formSection.classList.contains("hidden")) {
      return notMain;
    }
  } else if (notMain === "bookmark list") {
    if (mainSection.classList.contains("hidden") && !bookmarkListSection.classList.contains("hidden")) {
      return "main"
    } else if (!mainSection.classList.contains("hidden") && bookmarkListSection.classList.contains("hidden")) {
      return notMain;
    }
  }  
}

const originalLocalStorage = JSON.parse(localStorage.getItem("bookmarks")) || [];
const temp = originalLocalStorage;
function resetLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify(temp));
}

function setLocalStorage() {
  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}]));
}

function clearCategoryList() {
  const categoryList = document.getElementById("category-list");
  categoryList.innerHTML = "";
}
```

# --hints--

Debe tener una función `getBookmarks`.

```js
assert.isFunction(getBookmarks);
```

Su función `getBookmarks` debe devolver un array.

```js
assert.isArray(getBookmarks());
```

Su función `getBookmarks` debe devolver la clave `bookmarks` de `localStorage`.

```js
try {
  setLocalStorage()
  const testArr = getBookmarks();
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(testArr, expected);
} finally {
  resetLocalStorage();
}
```

Cuando la clave `bookmarks` no está configurada en `localStorage` o es un array vacío, la función `getBookmarks` debe devolver un array vacío.

```js
try {
  localStorage.setItem("bookmarks", JSON.stringify([]));
  assert.isEmpty(getBookmarks());
  localStorage.removeItem("bookmarks");
  assert.isEmpty(getBookmarks());
} finally {
  resetLocalStorage();
}
```

Cuando la clave `bookmarks` en `localStorage` no contiene un array válido de objetos de marcador, la función `getBookmarks` debe devolver un array vacío.

```js
try {
  localStorage.setItem("bookmarks", 'invalid');
  const arrayFromInvalidValue = getBookmarks();
  assert.isArray(arrayFromInvalidValue);
  assert.isEmpty(arrayFromInvalidValue);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "example1", category: "news", notUrl: "example1.com"}]));
  const arrayFromInvalidArray = getBookmarks();
  assert.isArray(arrayFromInvalidArray)
  assert.isEmpty(arrayFromInvalidArray)
} finally {
  resetLocalStorage();
}
```

Debe tener una función llamada `displayOrCloseForm`.

```js
assert.isFunction(displayOrCloseForm);
```

Su función `displayOrCloseForm` debe alternar la clase `hidden` en `#main-section` y `#form-section`.

```js
const hidden = getHidden("form");
displayOrCloseForm();
const hiddenAfterCall = getHidden("form");
assert.isDefined(hidden, hiddenAfterCall);
assert.notEqual(hidden, hiddenAfterCall);
```

Al hacer clic en `#add-bookmark-button`, debe actualizar el texto interno de `.category-name` al valor de la opción seleccionada de `#category-dropdown`.

```js
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
addBookMarkButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText.trim();
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

Al hacer clic en `#add-bookmark-button`, debe llamar a `displayOrCloseForm` para mostrar la sección del formulario y ocultar la sección principal.

```js
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
assert.strictEqual(getHidden("form"), "form");
addBookMarkButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "main");
```

Al hacer clic en `#close-form-button`, debe llamar a `displayOrCloseForm` para ocultar la sección del formulario y mostrar la sección principal.

```js
const closeFormButtonTest = document.getElementById("close-form-button");
assert.strictEqual(getHidden("form"), "main");
closeFormButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("form"), "form");
```

Al hacer clic en `#add-bookmark-button-form`, debe actualizar la clave `bookmarks` almacenada en el almacenamiento local agregando un objeto al final del array. El objeto añadido debe tener `name` establecido en el valor del input `#name`, `category` establecido en el valor de la opción seleccionada del desplegable de categorías y `url` establecido en el valor del input `#url`.

```js
setLocalStorage();
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const categoryDropdownTest = document.getElementById("category-dropdown");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example5";
  bookmarkURLTest.value = "test example5.com";
  const bookmarkObject = {
    name: bookmarkNameTest.value,
    category: categoryDropdownTest.value,
    url: bookmarkURLTest.value
  };
  const expected = [{name: "example1", category: "news", url: "example1.com"}, {name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}, bookmarkObject];
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.deepEqual(getBookmarks(), expected);
} finally {
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
  resetLocalStorage();
}
```

Al hacer clic en `#add-bookmark-button-form`, debe restablecer el valor de `#name` y `#url` a una cadena vacía.

```js
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example2";
  bookmarkURLTest.value = "test example2.com";
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(bookmarkNameTest.value, "");
  assert.strictEqual(bookmarkURLTest.value, "");
} finally {
  resetLocalStorage()  
  bookmarkNameTest.value = "";
  bookmarkURLTest.value = "";
}
```

Al hacer clic en `#add-bookmark-button-form`, debe ejecutar `displayOrCloseForm` para ocultar la sección del formulario y mostrar la sección principal.

```js
const addBookmarkButtonFormTest = document.getElementById("add-bookmark-button-form");
const addBookMarkButtonTest = document.getElementById("add-bookmark-button");
const bookmarkNameTest = document.getElementById("name");
const bookmarkURLTest = document.getElementById("url");
try {
  bookmarkNameTest.value = "test example3";
  bookmarkURLTest.value = "test example3.com";  
  addBookMarkButtonTest.dispatchEvent(new Event("click"));  
  assert.strictEqual(getHidden("form"), "main");
  addBookmarkButtonFormTest.dispatchEvent(new Event("click"));
  assert.strictEqual(getHidden("form"), "form");
} finally {
  resetLocalStorage()  
}
```

Debe tener una función llamada `displayOrHideCategory`.

```js
assert.isFunction(displayOrHideCategory);
```

Su función `displayOrHideCategory` debe alternar la clase `hidden` en `#main-section` y `#bookmark-list-section`.

```js
const hidden = getHidden("bookmark list");
displayOrHideCategory();
const hiddenAfterCall = getHidden("bookmark list");
assert.isDefined(hidden, hiddenAfterCall)
assert.notEqual(hidden, hiddenAfterCall)
```

Al hacer clic en `#view-category-button`, debe actualizar el texto interno de `.category-name` al valor de la opción seleccionada de `#category-dropdown`.

```js
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
viewCategoryButtonTest.dispatchEvent(new Event("click"));
const categoryName = document.querySelector(".category-name").innerText.trim();
assert.strictEqual(categoryName.toLowerCase(), categoryDropdownTest.value.toLowerCase());
```

Al hacer clic en `#view-category-button`, debe agregar un elemento `p` con el texto `No se encontraron marcadores` en el HTML interno de `#category-list` si ninguno de los marcadores en el almacenamiento local tiene la categoría seleccionada.

```js
try {
  const viewCategoryButtonTest = document.getElementById("view-category-button");
  const categoryDropdownTest = document.getElementById("category-dropdown");
  setLocalStorage()
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const categoryList = document.getElementById("category-list").innerHTML;
  assert.strictEqual(categoryList, "<p>No Bookmarks Found</p>")
} finally {
  resetLocalStorage()
  clearCategoryList()
}
```

Al hacer clic en el `#view-category-button`, debe modificar el elemento `#category-list` en el HTML interno agregando un botón de radio. El botón de radio debe tener los atributos `id` y `value` establecidos en el nombre del marcador para cada marcador en la categoría seleccionada. Además, cada botón de radio debe tener el mismo atributo `name`.

```js
try {
  const viewCategoryButtonTest = document.getElementById("view-category-button");
  const categoryDropdownTest = document.getElementById("category-dropdown");
  setLocalStorage();
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(bookmarksDisplayed, 2);
  assert.strictEqual(bookmarksDisplayed[0].id, "example1");
  assert.strictEqual(bookmarksDisplayed[0].value, "example1");
  assert.strictEqual(bookmarksDisplayed[1].id, "example4");
  assert.strictEqual(bookmarksDisplayed[1].value, "example4");
  assert.isOk(bookmarksDisplayed[0].name)
  assert.strictEqual(bookmarksDisplayed[0].name, bookmarksDisplayed[1].name);
} finally {
  resetLocalStorage();
}
```

Cada botón de radio añadido al HTML interno de `#category-list` debe tener una etiqueta correspondiente que contenga un elemento de anclaje con el nombre del marcador y el atributo `href` establecido en la URL del marcador.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarkLabelsDisplayed, 2);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example1");
  assert.strictEqual(bookmarkLabelsDisplayed[1].htmlFor, "example4");
} finally {
  resetLocalStorage();
}
```

Cada elemento `label` debe contener un elemento de anclaje con el nombre del marcador como texto, y el atributo `href` establecido en la URL del marcador.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const bookmarkAnchorsDisplayed = document.querySelectorAll('#category-list label>a');
  assert.lengthOf(bookmarkAnchorsDisplayed, 2);
  assert.strictEqual(bookmarkAnchorsDisplayed[0].innerText.trim(), "example1");
  assert.isTrue(bookmarkAnchorsDisplayed[0].href.endsWith("example1.com"));
  assert.strictEqual(bookmarkAnchorsDisplayed[1].innerText.trim(), "example4");
  assert.isTrue(bookmarkAnchorsDisplayed[1].href.endsWith("example4.com"));
} finally {
  resetLocalStorage();
}
```

Al hacer clic en `#close-list-button`, debe ocultar `#bookmark-list-section` y mostrar la sección principal.

```js
const viewCategoryButtonTest = document.getElementById("view-category-button");
const closeListButtonTest = document.getElementById("close-list-button");
viewCategoryButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "main");
closeListButtonTest.dispatchEvent(new Event("click"));
assert.strictEqual(getHidden("bookmark list"), "bookmark list");
```

Al hacer clic en el `#close-list-button` y luego abrir cualquier categoría, el `#category-list` debe contener solo datos relevantes para la categoría seleccionada, sin duplicar entradas.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const closeListButtonTest = document.getElementById("close-list-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const categoryList = document.getElementById("category-list").innerHTML;
  assert.strictEqual(categoryList, "<p>No Bookmarks Found</p>");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "news";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const newsBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(newsBookmarks, 2);
  assert.strictEqual(newsBookmarks[0].id, "example1");
  assert.strictEqual(newsBookmarks[0].value, "example1");
  assert.strictEqual(newsBookmarks[1].id, "example4");
  assert.strictEqual(newsBookmarks[1].value, "example4");

  closeListButtonTest.dispatchEvent(new Event("click"));
  categoryDropdownTest.value = "work";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const workBookmarks = document.querySelectorAll('#category-list input[type="radio"]');
  assert.lengthOf(workBookmarks, 1);
  assert.strictEqual(workBookmarks[0].id, "example3");
  assert.strictEqual(workBookmarks[0].value, "example3");

  categoryDropdownTest.value = "miscellaneous";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));

  const paragraphs = document.querySelectorAll("#category-list p");
  assert.lengthOf(paragraphs, 1);
  assert.strictEqual(paragraphs[0].innerText.trim(), "No Bookmarks Found");
} finally {
  resetLocalStorage();
  clearCategoryList();
}
```

Al hacer clic en el `#delete-bookmark-button`, debe eliminar el marcador correspondiente al botón de radio seleccionado y la categoría adecuada del almacenamiento local y actualizar la lista de marcadores mostrada.

```js
setLocalStorage();
const viewCategoryButtonTest = document.getElementById("view-category-button");
const deleteBookmarkButtonTest = document.getElementById("delete-bookmark-button");
const categoryDropdownTest = document.getElementById("category-dropdown");
try {
  categoryDropdownTest.value = "news";  
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const firstRadioButton = document.querySelector('#category-list input[type="radio"]');
  firstRadioButton.checked = true;
  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));
  let bookmarksDisplayed = document.querySelectorAll('#category-list input[type="radio"]');
  let bookmarkLabelsDisplayed = document.querySelectorAll('#category-list label');
  assert.lengthOf(bookmarksDisplayed, 1);
  assert.strictEqual(bookmarksDisplayed[0].id, "example4");
  assert.strictEqual(bookmarksDisplayed[0].value, "example4");
  assert.lengthOf(bookmarkLabelsDisplayed, 1);
  assert.strictEqual(bookmarkLabelsDisplayed[0].htmlFor, "example4");
  const expected = [{name: "example2", category: "entertainment", url: "example2.com"}, {name: "example3", category: "work", url: "example3.com"}, {name: "example4", category: "news", url: "example4.com"}];
  assert.deepEqual(getBookmarks(), expected);

  localStorage.setItem("bookmarks", JSON.stringify([{name: "duplicated-name", category: "news", url: "example1.com"}, {name: "duplicated-name", category: "entertainment", url: "example2.com"}]));

  categoryDropdownTest.value = "entertainment";
  viewCategoryButtonTest.dispatchEvent(new Event("click"));
  const radioForDuplicate = document.querySelector('#category-list input[type="radio"]');
  radioForDuplicate.checked = true;

  deleteBookmarkButtonTest.dispatchEvent(new Event("click"));

  assert.deepEqual(getBookmarks(), [{name: "duplicated-name", category: "news", url: "example1.com"}])
} finally {
  resetLocalStorage();
}
```

# --seed--

## --seed-contents--

```js

```

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=0.0">
    <title>Bookmark Manager</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <section id="main-section">
            <h1>Bookmark Manager</h1>
            <div id="dropdown">
                <label for="category-dropdown">Select a category:</label>
                <select id="category-dropdown" name="options">
                    <option value="news" selected>News</option>
                    <option value="entertainment">Entertainment</option>
                    <option value="work">Work</option>
                    <option value="miscellaneous">Miscellaneous</option>
                </select>
            </div>
            <div id="buttons">
                <button type="button" id="view-category-button">View Category</button>
                <button type="button" id="add-bookmark-button">Add Bookmark</button>
            </div>
        </section>

        <section id="form-section" class="hidden">
            <form>
                <h2 class="category-name"></h2>
                <div>
                    <label for="name">Name:</label>
                    <input type="text" id="name">
                </div>
                <div>
                    <label for="url">URL:</label>
                    <input type="text" id="url">
                </div>
                <div>
                    <button type="button" id="close-form-button">Go Back</button>
                    <button type="button" id="add-bookmark-button-form">Add Bookmark</button>
                </div>
            </form>
        </section>

        <section id="bookmark-list-section" class="hidden">
            <h2 class="category-name"></h2>
            <div id="category-list">
            </div>
            <div>
                <button type="button" id="close-list-button">Go Back</button>
                <button type="button" id="delete-bookmark-button">Delete Bookmark</button>
            </div>
        </section>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-grey: #0a0a23;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

main {
  display: flex;
  justify-content: center;
}

body {
  background-color: var(--dark-grey);
}

.hidden {
  display: none;
}

section {
  display: flex;
  flex-direction: column;
  justify-content: center;
}

select,
input,
label {
  margin-left: 10px;
}

div {
  padding: 30px;
  display: flex;
  justify-content: center;
}

.close-form-button {
  background: none;
  border: none;
  cursor: pointer;
}

h1, h2 {
  margin-top: 20px;
  text-align: center;
}

#category-list {
  text-align: center;
  display: flex;
  flex-direction: column;
  justify-content: flex-start;
  background-color: var(--light-grey);
  align-self: center;
  width: 80%;
  margin-top: 15px;
  border-radius: 10px;
}

#category-list,
h1,
h2,
label {
  color: var(--light-grey);
}

#category-list p {
  color: var(--dark-grey);
}

button {
  cursor: pointer;
  padding: 5px;
  width: 100px;
  margin: 10px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

section {
  margin-top: 60px;
  border: 2px solid var(--golden-yellow);
  width: fit-content;
  border-radius: 10px;
}
```

# --solutions--

```js
const categoryDropdown = document.getElementById("category-dropdown");
const viewCategoryButton = document.getElementById("view-category-button");
const addBookmarkButton = document.getElementById("add-bookmark-button");
const mainSection = document.getElementById("main-section");
const formSection = document.getElementById("form-section");
const bookmarkListSection = document.getElementById("bookmark-list-section");
const categoryHeading = document.querySelectorAll(".category-name");
const closeFormButton = document.getElementById("close-form-button");
const bookmarkName = document.getElementById("name");
const bookmarkURL = document.getElementById("url");
const addBookmarkButtonForm = document.getElementById("add-bookmark-button-form");
const categoryList = document.getElementById("category-list");
const closeListButton = document.getElementById("close-list-button");
const deleteBookmarkButton = document.getElementById("delete-bookmark-button");

const getBookmarks = () => {
  try {
    const parsed = JSON.parse(localStorage.getItem("bookmarks"));
    if (
      Array.isArray(parsed) &&
      parsed.every(
        (item) =>
          item.hasOwnProperty("category") &&
          item.hasOwnProperty("name") &&
          item.hasOwnProperty("url"),
      )
    ) {
      return parsed;
    }
    return [];
  } catch {
    return [];
  }
};

const displayOrCloseForm = () => {
    mainSection.classList.toggle("hidden");
    formSection.classList.toggle("hidden");
}

addBookmarkButton.addEventListener("click", () => {
    categoryHeading[0].innerText = categoryDropdown.value.charAt(0).toUpperCase()
        + categoryDropdown.value.slice(1);
    displayOrCloseForm();
});

closeFormButton.addEventListener("click", displayOrCloseForm);

const addBookmark = () => {
    const bookmarkObject = {
        name: bookmarkName.value,
        category: categoryDropdown.value,
        url: bookmarkURL.value
    };
    if (bookmarkName.value.length && bookmarkURL.value.length) {        
        localStorage.setItem("bookmarks", JSON.stringify(getBookmarks().concat(bookmarkObject)))
        bookmarkName.value = "";
        bookmarkURL.value = "";
        displayOrCloseForm()
    } else {
        alert("Please, provide valid name and URL.")
    }
}

addBookmarkButtonForm.addEventListener("click", addBookmark);

const displayOrHideCategory = () => {
    mainSection.classList.toggle("hidden");
    bookmarkListSection.classList.toggle("hidden");
}

const fillBookmarkList = () => {
    categoryHeading[1].innerText = categoryDropdown.value.charAt(0).toUpperCase()
        + categoryDropdown.value.slice(1);
    const bookmarksToDisplay = getBookmarks().filter((i) => i.category === categoryDropdown.value);
    if (bookmarksToDisplay.length) {
        categoryList.innerHTML = '';
        for (const bookmark of bookmarksToDisplay) {
            categoryList.innerHTML += `
            <div>
              <input type="radio" name="bookmarks" id="${bookmark.name}" value="${bookmark.name}">
              <label for="${bookmark.name}">
                <a href="${bookmark.url}">${bookmark.name}</a>
              </label>
            </div>`
        }
    } else {
        categoryList.innerHTML = '<p>No Bookmarks Found</p>';
    }
}

viewCategoryButton.addEventListener("click", () => {
    fillBookmarkList();
    displayOrHideCategory();
})

closeListButton.addEventListener("click", displayOrHideCategory)

const deleteBookmark = () => {
    const radioBookmarks = document.querySelectorAll('input[type="radio"]');
    for (const radioBookmark of radioBookmarks) {
        if (radioBookmark.checked) {
            const bookmarks = getBookmarks();
            const indexToRemove = bookmarks.findIndex(i => i.name == radioBookmark.value && i.category === categoryDropdown.value);
            bookmarks.splice(indexToRemove, 1);
            localStorage.setItem("bookmarks", JSON.stringify(bookmarks));
            return
        }
    }
}

deleteBookmarkButton.addEventListener("click", () => {
    deleteBookmark();
    fillBookmarkList();
})
```

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=0.0">
    <title>Bookmark Manager</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <section id="main-section">
            <h1>Bookmark Manager</h1>
            <div id="dropdown">
                <label for="category-dropdown">Select a category:</label>
                <select id="category-dropdown" name="options">
                    <option value="news" selected>News</option>
                    <option value="entertainment">Entertainment</option>
                    <option value="work">Work</option>
                    <option value="miscellaneous">Miscellaneous</option>
                </select>
            </div>
            <div id="buttons">
                <button type="button" id="view-category-button">View Category</button>
                <button type="button" id="add-bookmark-button">Add Bookmark</button>
            </div>
        </section>

        <section id="form-section" class="hidden">
            <form>
                <h2 class="category-name"></h2>
                <div>
                    <label for="name">Name:</label>
                    <input type="text" id="name">
                </div>
                <div>
                    <label for="url">URL:</label>
                    <input type="text" id="url">
                </div>
                <div>
                    <button type="button" id="close-form-button">Go Back</button>
                    <button type="button" id="add-bookmark-button-form">Add Bookmark</button>
                </div>
            </form>
        </section>

        <section id="bookmark-list-section" class="hidden">
            <h2 class="category-name"></h2>
            <div id="category-list">
            </div>
            <div>
                <button type="button" id="close-list-button">Go Back</button>
                <button type="button" id="delete-bookmark-button">Delete Bookmark</button>
            </div>
        </section>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-grey: #0a0a23;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

main {
  display: flex;
  justify-content: center;
}

body {
  background-color: var(--dark-grey);
}

.hidden {
  display: none;
}

section {
  display: flex;
  flex-direction: column;
  justify-content: center;
}

select,
input,
label {
  margin-left: 10px;
}

div {
  padding: 30px;
  display: flex;
  justify-content: center;
}

.close-form-button {
  background: none;
  border: none;
  cursor: pointer;
}

h1, h2 {
  margin-top: 20px;
  text-align: center;
}

#category-list {
  text-align: center;
  display: flex;
  flex-direction: column;
  justify-content: flex-start;
  background-color: var(--light-grey);
  align-self: center;
  width: 80%;
  margin-top: 15px;
  border-radius: 10px;
}

#category-list,
h1,
h2,
label {
  color: var(--light-grey);
}

#category-list p {
  color: var(--dark-grey);
}

button {
  cursor: pointer;
  padding: 5px;
  width: 100px;
  margin: 10px;
  color: var(--dark-grey);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

section {
  margin-top: 60px;
  border: 2px solid var(--golden-yellow);
  width: fit-content;
  border-radius: 10px;
}
```
