---
id: 6a57c331ae94fa4631b60864
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

审计从读取已有的记录开始。由于每个工件都保存在其 ID 下，你可以通过将该 ID 放入方括弧中来访问它，比如 `collection[101]`。

调用 `console.log()` 并传入 `collection[101].title` 来记录第一个工件的标题。然后，调用 `console.log()` 并传入 `collection[101].curator.name` 来记录其策展人的名字。

# --hints--

你应该将第一个工件的 `title` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)/);
```

你应该将第一个工件策展人的 `name` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

你应该先记录 `title`，然后记录策展人的 `name`。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*title|\[\s*["']title["']\s*\])\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

--fcc-editable-region--

--fcc-editable-region--
```
