---
id: 6a57c265abbc95dfebc05098
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将构建一个文物来源审计器。你会将博物馆文物保存在一个对象中，并编写函数来查找它们、为它们添加标签、在展厅之间移动它们，以及总结它们的历史。

每件文物都以其 ID 进行保存。创建一个名为 `collection` 的对象并向其中添加一个条目。将密钥设置为 `101`，值设置为一个对象，该对象的 `title` 为 `"Golden Mask"`，`category` 为 `"Ceremonial"`。

# --hints--

你应该声明一个名为 `collection` 的变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` 应该是一个对象。

```js
assert.isObject(collection);
```

`collection` 应该有一个条目。

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

存储在 `101` 下的文物应该是一个对象。

```js
assert.isObject(collection[101]);
```

存储在 `101` 下的文物应该有一个 `title`，值为 `"Golden Mask"`。

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

存储在 `101` 下的文物应该有一个 `category`，值为 `"Ceremonial"`。

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
