---
id: 6a57c331ae94fa4631b60868
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

策展人会为文物添加标签以便更容易查找，且标签绝不应重复添加。在添加新标签之前，你可以使用 `includes()` 方法来查看它是否已经存在，因为当数组已经包含某个值时，该方法会返回 `true`。

创建一个名为 `addTag` 的函数，接收 `id` 和 `tag` 两个参数。使用 `collection[id]` 查找文物。如果文物存在且其 `tags` 数组中还没有包含 `tag`，则将 `tag` 推入 `artifact.tags`。

# --hints--

你应该有一个名为 `addTag` 的函数，参数名为 `id` 和 `tag`。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

调用 `addTag` 并传入一个新标签时，应将该标签添加到文物的 `tags` 数组中。

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

调用 `addTag` 并传入文物已有的标签时，不应添加重复标签。

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` 应该将标签添加到 `id` 所指向的任何文物上。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

当没有文物匹配 `id` 时，`addTag` 不应抛出误差或更改集合。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

你应该将 `tag` 推入文物现有的 `tags` 数组，而不是替换整个数组。

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
