---
id: 6a57c331ae94fa4631b60870
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

最后，为任何工件构建一个可读的摘要。首先处理工件缺失的分支。

创建一个名为 `buildSummary` 的函数，它接受一个 `id` 参数。使用 `collection[id]` 查找工件，并将其保存到名为 `artifact` 的变量中。如果未找到工件，返回字符串 `"Artifact not found"`。

# --hints--

你应该有一个名为 `buildSummary` 的函数，带有一个名为 `id` 的参数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary);
const params = explorer.allFunctions.buildSummary.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

你应该将 `collection[id]` 赋值给一个名为 `artifact` 的变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.artifact);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.artifact?.value.matches("collection[id]"));
```

当没有工件匹配 `id` 时，`buildSummary` 应该返回 `"Artifact not found"`。

```js
assert.strictEqual(buildSummary(999), "Artifact not found");
```

当工件存在时，`buildSummary` 不应该返回 `"Artifact not found"`。

```js
assert.notStrictEqual(buildSummary(101), "Artifact not found");
assert.notStrictEqual(buildSummary(102), "Artifact not found");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
