---
id: 6a57c331ae94fa4631b6086b
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

将 `id` 为 `102` 的文物在 `2026` 年移至 `"Hall B"`。然后，调用 `console.log()` 并传入 `collection[102].locations` 来记录该文物的 locations 数组。

# --hints--

你应该用 `102`、`"Hall B"` 和 `2026` 调用 `moveArtifact`。

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)/);
```

你应该将第二个文物的 `locations` 记录到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

你应该在移动文物后记录 `locations`。

```js
assert.match(__helpers.removeJSComments(code), /moveArtifact\s*\(\s*102\s*,\s*["']Hall B["']\s*,\s*2026\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?102["']?\s*\](?:\.\s*locations|\[\s*["']locations["']\s*\])\s*\)/);
```

第二个文物的最新位置应该是 `{ gallery: "Hall B", year: 2026 }`。

```js
assert.deepEqual(collection[102].locations[collection[102].locations.length - 1], { gallery: "Hall B", year: 2026 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
