---
id: 6a57c331ae94fa4631b60861
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

每个文物都有一个负责人。

为存储在密钥 `101` 下的文物添加一个 `curator` 属性，并将其值设置为一个对象，该对象有一个 `id` 为 `201` 和一个 `name` 为 `"Earl Sinclair"`。

# --hints--

存储在 `101` 下的文物应该有一个 `curator` 属性。

```js
assert.property(collection[101], "curator");
```

`curator` 应该是一个对象。

```js
assert.isObject(collection[101].curator);
```

`curator` 应该有一个 `id` 为 `201`。

```js
assert.strictEqual(collection[101].curator.id, 201);
```

`curator` 应该有一个 `name` 为 `"Earl Sinclair"`。

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
