---
id: 6a57c331ae94fa4631b60871
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

一个文物当前的展馆是它最近的位置。你可以通过查看它的 `locations` 数组中的最后一项来找到这个位置。

声明一个 `currentLocation` 变量，并赋值为 `artifact.locations[artifact.locations.length - 1]`。这个表达式使用数组的 `length` 属性来访问它的最后一个元素。

# --hints--

你应该在 `buildSummary` 内声明一个名为 `currentLocation` 的变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.buildSummary?.variables.currentLocation);
```

你应该将文物的 `locations` 数组的最后一个元素赋值给 `currentLocation`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.buildSummary?.variables.currentLocation?.value.matches("artifact.locations[artifact.locations.length - 1]"));
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

updateCurator(101, "Fran Sinclair");
console.log(collection[101].curator.name);

function buildSummary(id) {
  const artifact = collection[id];
  if (!artifact) {
    return "Artifact not found";
  }
--fcc-editable-region--
  
--fcc-editable-region--
}
```
