---
id: 6a57c331ae94fa4631b60866
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

有时工作人员会要求博物馆没有的文物。如果在该 `id` 下没有保存任何文物，`collection[id]` 就是 `undefined`，读取它的 `.title` 会抛出误差。

更新返回语句以使用条件式。如果找到文物，返回 `artifact.title`，如果没有找到，则返回字串 `"Artifact not found"`。

# --hints--

`getArtifactTitle(101)` 仍应返回 `"Golden Mask"`。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

如果没有文物匹配该 `id`，`getArtifactTitle` 应返回 `"Artifact not found"`。

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` 应返回它找到的任何文物的标题，而不仅仅是 `101`。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
