---
id: 6a57c331ae94fa4631b60865
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

工作人员每天要查找几十次文物，所以你的审计员需要一个可重用的查找方法。方括弧也可以接受变量，因此 `collection[id]` 会返回 `id` 所参考的那个文物。

创建一个名为 `getArtifactTitle` 的函数，它接受一个 `id` 参数。使用 `collection[id]` 获取文物，将其保存到名为 `artifact` 的变量中，并返回 `artifact.title`。

# --hints--

你应该有一个名为 `getArtifactTitle` 的函数，并带有一个名为 `id` 的参数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

你应该将 `collection[id]` 赋值给一个名为 `artifact` 的变量。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` 应该返回 `artifact.title`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` 应该返回 `"Golden Mask"`。

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` 应该返回 `"Bronze Tablet"`。

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
