---
id: 6a57c331ae94fa4631b60862
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

每个文物还会记录它曾经所在的位置、它的描述方式，以及它是否正在展出。

为存储在 `101` 下的文物添加三个属性：

- `locations`：一个包含两个对象的数组，分别是 `{ gallery: "Hall A", year: 2020 }` 和 `{ gallery: "Hall C", year: 2024 }`
- `tags`：一个包含字串 `"gold"` 和 `"egypt"` 的数组
- `onDisplay`：布尔值 `true`

# --hints--

文物的 `locations` 应该是包含这两个给定位置对象的数组。

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

文物的 `tags` 应该是 `["gold", "egypt"]`。

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

文物的 `onDisplay` 应该是 `true`。

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
