---
id: 6a57c331ae94fa4631b6086e
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

随着你的博物馆不断发展，负责某件文物的人可能会发生变化。

创建一个名为 `updateCurator` 的函数，它接受 `id` 和 `name` 作为参数。使用 `collection[id]` 查找文物。如果存在，则将其 `curator` 对象的 `name` 属性更新为提供的 `name`。

# --hints--

你应该有一个名为 `updateCurator` 的函数，参数名为 `id` 和 `name`。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.updateCurator);
const params = explorer.allFunctions.updateCurator.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "name");
```

调用 `updateCurator` 应该会更改文物的 `curator` 对象的 `name`。

```js
const curatorBefore = { ...collection[102].curator };
updateCurator(102, "Test Curator");
assert.strictEqual(collection[102].curator.name, "Test Curator");
assert.deepEqual(collection[102].curator, {
  ...curatorBefore,
  name: "Test Curator",
});
```

`updateCurator` 应该更改 `id` 所指向的任何文物的馆长。

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
updateCurator(457, "Fresh Curator");
assert.strictEqual(collection[457].curator.name, "Fresh Curator");
```

当没有文物匹配该 `id` 时，`updateCurator` 不应抛出误差或更改集合。

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => updateCurator(999, "Test Curator"));
assert.strictEqual(JSON.stringify(collection), before);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

--fcc-editable-region--

--fcc-editable-region--
```
