---
id: 69de4308e652aaeb5ccdb03e
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nesta oficina, você vai praticar o uso dos métodos de string do Python.

Comece criando uma variável chamada `email` e atribua a ela a string `alice.johnson@company.com`.

# --hints--

Você deve criar uma variável chamada `email`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Você deve atribuir a string `alice.johnson@company.com` à variável `email`. Lembre-se de colocar o texto entre aspas simples ou aspas duplas.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
