---
id: 69e20876baed33de636b989a
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Às vezes, você pode querer buscar a posição de um caractere específico dentro de uma string. O Python oferece o método embutido `.find()` para esse propósito. Ele retorna um inteiro representando o índice (posição) da primeira ocorrência da substring dada.

Aqui está um exemplo:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Agora, faça a chamada do método `.find()` em `email` e passe `'@'` para buscar a posição do sinal de arroba na string `email`. Atribua o resultado a uma variável chamada `at_position`.

Depois, imprima `at_position` para que a saída seja: `Position of @: 13`.

# --hints--

Você deve criar uma variável chamada `at_position`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Você deve atribuir a `at_position` o valor `email.find('@')`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Você deve imprimir `at_position` para que a saída seja `Position of @: 13`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
