---
id: 69e20a0451c7a3ef189514d3
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Agora que você sabe a posição do `@` no e-mail, pode extrair o nome de usuário, que fica antes do sinal `@`.

Na oficina anterior, você extraiu uma parte de uma string usando slicing. Para lembrar, a sintaxe é:

```python
text = 'hello world'
print(text[:5])  # hello
```

Use slicing no `email` para extrair tudo antes de `at_position` e armazene o resultado em uma variável chamada `username`.

Depois, imprima `username` para que a saída seja `Username: alice.johnson`.

# --hints--

Você deve criar uma variável chamada `username`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Você deve usar slicing com `at_position` no `email` para atribuir um valor a `username`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Você deve imprimir `username` para que a saída seja `Username: alice.johnson`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
