---
id: 69e6299710319e7cb8ebdc37
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Dados do mundo real costumam ser bagunçados, com espaços extras e uso inconsistente de maiúsculas e minúsculas.

Crie uma variável `messy_name` e atribua a ela a string `sARaH dAVis` com dois espaços no início e dois no final. Depois, imprima para que a saída seja:

```md
Original name: '  sARaH dAVis  '
```

Você vai limpar isso nos próximos passos.

# --hints--

Você deve criar uma variável chamada `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Você deve atribuir <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> a `messy_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Você deve imprimir `messy_name` para que a saída seja `Original name: '  sARaH dAVis  '`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
