---
id: 69e6284f3b7659c5a05c2a79
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Você pode verificar se uma substring existe dentro de uma string usando a palavra-chave `in`. Ela retorna `True` se a substring existir e `False` caso contrário. Aqui está um exemplo:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Crie uma variável chamada `has_at` que verifica se o caractere `@` está no `email`.

Depois, imprima `has_at` para que a saída seja: `Has @ symbol: True`.

# --hints--

Você deve criar uma variável chamada `has_at`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Sua variável `has_at` deve ter o valor `'@' in email`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Você deve imprimir `has_at` para que a saída seja `Has @ symbol: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
