---
id: 69e7766b93d27261775d1cf1
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Agora imprima `employee_code`, `starts_with_emp`, `clean_email` e `ends_with_com` para que a saída seja:

```md
Code: EMP-2024-SD
Starts with EMP: True
Clean email: sarah.davis@company.com
Ends with .com: True
```

# --hints--

Você deve imprimir `employee_code` para que a saída seja `Code: EMP-2024-SD`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Code: {employee_code}')",
    "print('Code: ' + employee_code)",
    "print('Code:', employee_code)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Você deve imprimir `starts_with_emp` para que a saída seja `Starts with EMP: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Starts with EMP: {starts_with_emp}')",
    "print('Starts with EMP: ' + str(starts_with_emp))",
    "print('Starts with EMP:', starts_with_emp)",
    "print('Starts with EMP:', str(starts_with_emp))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Você deve imprimir `clean_email` para que a saída seja `Clean email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Clean email: {clean_email}')",
    "print('Clean email: ' + clean_email)",
    "print('Clean email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Você deve imprimir `ends_with_com` para que a saída seja `Ends with .com: True`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Ends with .com: {ends_with_com}')",
    "print('Ends with .com: ' + str(ends_with_com))",
    "print('Ends with .com:', ends_with_com)",
    "print('Ends with .com:', str(ends_with_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

--fcc-editable-region--

--fcc-editable-region--
```
