---
id: 69e73f4f5f986657647bbd63
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

O método `.replace(old, new)` substitui toda ocorrência da substring `old` por outra, `new`. Você também pode encadeá-lo com outros métodos. Por exemplo:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

Substitua `.` por um único espaço em `email_user` e depois chame `.title()` no resultado. Atribua o resultado a uma variável chamada `display_name`.

Depois imprima `display_name` para que a saída seja: `Display name: Sarah Davis`.

# --hints--

Você deve criar uma variável chamada `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

Você deve usar `.replace('.', ' ')` e `.title()` em `email_user` e atribuir o resultado à variável `display_name`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

Você deve imprimir `display_name` para que a saída seja `Display name: Sarah Davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
