---
id: 69e73a141addc6f442924b0d
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Você já usou `.find()` e slicing na variável `email`. Agora vai usar a mesma técnica em `clean_email`.

Busque a posição do caractere `@` em `clean_email` e armazene em `at_pos`. Depois use slicing para extrair tudo antes do `@` e armazene em `email_user`.

Depois imprima `clean_email` e `email_user` para que a saída seja:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Você deve criar uma variável chamada `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Você deve chamar `.find('@')` em `clean_email` e atribuir o resultado à variável `at_pos`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Você deve criar uma variável chamada `email_user`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Você deve usar slicing com `at_pos` em `clean_email` e atribuir o resultado à variável `email_user`.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Você deve imprimir `clean_email` para que a saída seja `Email: sarah.davis@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Você deve imprimir `email_user` para que a saída seja `Username: sarah.davis`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
