---
id: 69e76d1997661be7f0040e2d
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

O método `.split(separator)` divide uma string em uma lista de substrings a cada ocorrência do `separator`. Por exemplo:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

Crie uma variável `full_address` e atribua a ela a string `123 Main Street, Springfield, IL`. Depois crie `address_parts` dividindo pela vírgula seguida de um espaço.

Depois imprima `address_parts` para que a saída seja `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

# --hints--

Você deve criar uma variável chamada `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

Você deve atribuir a string `123 Main Street, Springfield, IL` a `full_address`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

Você deve criar uma variável chamada `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

Você deve usar `.split(', ')` em `full_address` e atribuir o resultado à variável `address_parts`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

Você deve imprimir `address_parts` para que a saída seja `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
