---
id: 69e774b82fe0bfdb684f95b4
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Às vezes você precisa verificar como uma string começa. Python oferece um método `.startswith()` para isso. Ele retorna `True` se a string começar com um dado prefixo; caso contrário, retorna `False`.

Aqui está um exemplo:

```python
code = 'EMP-001'
print(code.startswith('EMP')) # True
print(code.startswith('001')) # False
```

Crie uma variável `employee_code` e atribua a ela a string `EMP-2024-SD`. Depois crie uma variável `starts_with_emp` que verifica se `employee_code` começa com `EMP`.

# --hints--

Você deve criar uma variável chamada `employee_code`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('employee_code')
`) })
```

Você deve atribuir `EMP-2024-SD` a `employee_code`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('employee_code').is_equivalent("employee_code = 'EMP-2024-SD'")
`) })
```

Você deve criar uma variável chamada `starts_with_emp`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('starts_with_emp')
`) })
```

Você deve usar `.startswith('EMP')` em `employee_code` e atribuir o resultado à variável `starts_with_emp`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('starts_with_emp').is_equivalent("starts_with_emp = employee_code.startswith('EMP')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

--fcc-editable-region--

--fcc-editable-region--
```
