---
id: 6a86be3e8cc29e9ae239e70a
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Você também pode precisar verificar como uma string termina. Python oferece o método `.endswith()` para isso. Ele retorna `True` se a string termina com o sufixo dado e `False` se a string não termina com o sufixo dado.

Aqui está um exemplo:

```python
code = 'EMP-001'
print(code.endswith('EMP')) # False
print(code.endswith('001')) # True
```

Crie uma variável `ends_with_com` que verifica se `clean_email` termina com `.com`.

# --hints--

Você deve criar uma variável chamada `ends_with_com`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('ends_with_com')
`) })
```

Você deve usar `.endswith('.com')` em `clean_email` e atribuir o resultado à variável `ends_with_com`.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('ends_with_com').is_equivalent("ends_with_com = clean_email.endswith('.com')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')

--fcc-editable-region--

--fcc-editable-region--
```
