---
id: 6a870cc60c23f0332a2fe788
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Lembre-se de uma oficina anterior que você pode usar uma f-string para incluir uma variável dentro de uma string, escrevendo um `f` antes da aspa de abertura e envolvendo a variável em `{}`.

Aqui está um lembrete de como funciona:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Agora, imprima `email` para que a saída seja: `Email: alice.johnson@company.com`.

# --hints--

Você deve imprimir `email` para que a saída seja `Email: alice.johnson@company.com`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
