---
id: 67116c175a77cb64c5c10d49
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이전 수업에서 `map`, `filter`, `reduce` 같은 고차 함수를 다루는 방법을 배웠습니다.

이번 워크숍에서는 다양한 고차 함수를 다루는 연습을 할 수 있는 도서관 관리자 앱을 만들어 보세요.

먼저, `library`라는 변수를 만들고 빈 배열로 할당하세요.

# --hints--

`library` 배열을 만들어야 합니다.

```js
assert.isArray(library)
```

`library` 배열은 비어 있어야 합니다.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
