---
id: 67116d7584d0b469b14579c0
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

다음 속성과 값을 가진 객체를 `library` 배열 안에 또 하나 만드세요:

| 속성 | 값   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


`library` 배열의 두 번째 항목으로 객체가 있어야 합니다.

```js
assert.isObject(library[1]);
```

객체에 `title` 속성이 있어야 합니다.

```js
assert.property(library[1], "title")
```

`title` 속성의 값은 `"Atomic Habits"` 여야 합니다.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

객체에 `author` 속성이 있어야 합니다.

```js
assert.property(library[1], "author")
```

`author` 속성의 값은 `"James Clear"` 여야 합니다.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

객체에 `about` 속성이 있어야 합니다.

```js
assert.property(library[1], "about")
```

`about` 속성의 값은 `"A practical book about discarding bad habits and building good ones"` 이어야 합니다.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

객체에 `pages` 속성이 있어야 합니다.

```js
assert.property(library[1], "pages")
```

`pages` 속성의 값은 숫자여야 합니다.

```js
assert.isNumber(library[1]?.pages)
```

`pages` 속성은 `320` 로 설정해야 합니다.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
