---
id: 67116d7584d0b469b14579bf
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

`library` 배열 안에 다음 속성과 값을 가진 객체를 만드세요:

| 속성 | 값   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

`library` 배열 안에 객체가 있어야 합니다.

```js
assert.isObject(library[0])
```

객체에 `title` 속성이 있어야 합니다.

```js
assert.property(library[0], "title")
```

`title` 속성의 값은 `"Your Next Five Moves: Master the Art of Business Strategy"`여야 합니다.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

객체에 `author` 속성이 있어야 합니다.

```js
assert.property(library[0], "author")
```

`author` 속성의 값은 `"Patrick Bet-David and Greg Dinkin"`여야 합니다.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

객체에 `about` 속성이 있어야 합니다.

```js
assert.property(library[0], "about")
```

`about` 속성의 값은 `"A book on how to plan ahead"`여야 합니다.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

객체에 `pages` 속성이 있어야 합니다.

```js
assert.property(library[0], "pages")
```

`pages` 속성의 값은 숫자여야 합니다.

```js
assert.isNumber(library[0]?.pages)
```

`pages` 속성은 `320`로 설정해야 합니다.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
