---
id: 681dc623b2b18887266777b1
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이전 수업에서, 주어진 함수를 원래 배열의 각 요소에 적용하여 새 배열을 만드는 `map()` 메서드에 대해 배웠습니다.

예시는 다음과 같습니다:

```js
const developers = [
  { name: "Alice", city: "New York", title: "Front-End Developer" },
  { name: "Bob", city: "San Francisco", title: "Back-End Developer" }
];

console.log(developers.map(dev => dev.name));
// ["Alice", "Bob"]
```

이 예시에서 `dev`은 `developers` 배열의 각 객체를 나타냅니다. 그런 다음 점 표기법을 사용해 객체에서 `name`를 가져옵니다. 마지막으로 결과는 이름들로 이루어진 새 배열이 됩니다.

`getBookInformation` 함수 안에서 `map()` 매개변수에 `catalog` 메서드를 사용해 책 제목만 포함된 새 배열을 반환하세요. 도움이 필요하면 예시를 참고하세요.

# --hints--

`getBookInformation` 함수는 배열을 반환해야 합니다.

```js
assert.isArray(getBookInformation(library));
```

`getBookInformation` 함수는 책 제목을 나타내는 문자열 배열을 반환해야 합니다.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = ["Title A", "Title B", "Title C"];

assert.deepEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  
}
--fcc-editable-region--
```
