---
id: 68225ea2f1aae301fa4df0f4
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

배열 결과가 작동하는 동안, 최종 원하는 결과는 제목과 저자가 있는 모든 책을 나열하는 문자열이어야 합니다.

여기서 `join()` 메서드가 필요합니다. 이전 수업에서 `join()` 메서드는 배열의 모든 요소를 지정한 구분자로 연결하여 하나의 새로운 문자열을 반환하는 데 사용된다는 것을 배웠습니다.

다시 한 번 살펴보겠습니다:

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

이 경우 구분자는 줄 바꿈 문자를 나타내는 `\n`입니다.

`join()`에 `\n` 메서드를 `map()` 구분자와 함께 연결하세요. 이제 결과 배열 대신 콘솔에서 문자열을 볼 수 있어야 합니다.

# --hints--

`getBookInformation` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(getBookInformation(library));
```

`getBookInformation` 함수는 최종 결과를 나타내는 하나의 문자열을 반환해야 합니다.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

--fcc-editable-region--
function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`)
}
--fcc-editable-region--

console.log(getBookInformation(library));
```
