---
id: 6822694e3d34c807feeaba50
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이전 단계와 마찬가지로, 원하는 결과는 배열 대신 문자열로 나타내는 것이 더 깔끔합니다.

기존 `join()` 메서드에 `map()` 메서드를 연결하고, 구분자로 `"\n"`를 `join()`에 전달하세요.

그런 다음 `getBookSummaries` 함수 아래에 `console.log()`를 추가하세요. `console` 문장은 `getBookSummaries` 인수를 사용해 `library` 함수를 호출해야 합니다.

이제 콘솔에서 결과가 문자열로 표시되는 것을 확인할 수 있습니다.

# --hints--

`getBookSummaries` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(getBookSummaries(library));
```

`getBookSummaries` 함수는 최종 결과를 나타내는 단일 문자열을 반환해야 합니다.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `About A
About B
About C`;

assert.strictEqual(getBookSummaries(testLibrary), expected);
```

코드에 `console.log(getBookSummaries(library));`이 있어야 합니다.

```js
assert.match(code, /console\.log\(\s*getBookSummaries\(\s*library\s*\)\);?/gi);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\nList of book summaries:\n");

--fcc-editable-region--
function getBookSummaries(catalog) {
  return catalog.map((book) => book.about)
}

--fcc-editable-region--
```
