---
id: 6827deed7fc01f074c90fc0d
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

다른 저자에 대해 `getBooksByAuthor` 함수를 테스트해보면 좋겠습니다.

먼저 `"\nList of books by James Clear:\n"` 메시지를 콘솔에 기록하세요.

그 아래에 `console.log()`를 추가하고, 또 다른 `console.log()`를 추가하세요. 그 `console.log()` 안에서 `getBooksByAuthor()` 함수를 `library`와 `"James Clear"` 인수로 호출하세요.

이제 콘솔에서 해당 저자의 모든 책을 볼 수 있을 것입니다.

# --hints--

`"\nList of books by James Clear:\n"` 메시지를 콘솔에 기록해야 합니다.

```js
assert.match(code, /console\.log\(("|')\\nList\s+of\s+books\s+by\s+James\s+Clear:\\n\1\);?/);
```

`console.log(getBooksByAuthor(library, "James Clear"))`가 있어야 합니다.

```js
assert.match(code, /console\.log\(\s*getBooksByAuthor\s*\(\s*library\s*,\s*(['"])\s*James\s+Clear\s*\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
  return catalog.map(book => `${book.title} by ${book.author}`).join("\n");
}

console.log(getBookInformation(library));

console.log("\nList of book summaries:\n");

function getBookSummaries(catalog) {
  return catalog.map((book) => book.about).join("\n");
}

console.log(getBookSummaries(library));

console.log("\nList of books by Arvid Kahl:\n");

function getBooksByAuthor(catalog, author) {
  return catalog.filter((book) => book.author === author);
}

console.log(getBooksByAuthor(library, "Arvid Kahl"));

--fcc-editable-region--

--fcc-editable-region--
```
