---
id: 66ec4c8e9878d8441956516f
title: 책 목록 표 만들기
challengeType: 25
dashedName: build-a-book-catalog-table
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 책 정보를 나열하는 `table` 요소를 만들어야 합니다.
1. 표에는 한 개의 행이 있는 헤더 요소가 있어야 합니다.
1. 표 헤더 요소의 행에는 네 개의 표 헤더 요소가 있어야 하며, 텍스트는 순서대로 `Title`, `Author`, `Genre`, `Publication Year`여야 합니다.
1. 표에는 최소 다섯 개의 행이 있는 본문 요소가 있어야 합니다.
1. 표 본문의 각 행에는 책의 제목, 저자, 장르, 출판 연도를 표시하는 네 개의 표 데이터 요소가 있어야 합니다.
1. 표에는 한 개의 행이 있는 푸터 요소가 있어야 합니다.
1. 푸터 요소의 행에는 네 개의 열을 병합한 표 데이터 요소가 있어야 하며, 텍스트는 `Total Books: N`이고 `N`는 표에 있는 책 수로 교체되어야 합니다.

# --hints--

하나의 `table` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

`thead` 요소 안에 하나의 `table` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table thead'), 1);
```

`tr` 요소 안에 하나의 `thead` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('thead tr'), 1);
```

`th` 요소의 행 안에 네 개의 `thead` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('thead tr th'), 4);
```

네 개의 `th` 요소는 순서대로 `Title`, `Author`, `Genre`, `Publication Year` 텍스트를 가져야 합니다.

```js
const ths = document.querySelectorAll('thead tr th');
assert.equal(ths[0]?.textContent.trim(), 'Title');
assert.equal(ths[1]?.textContent.trim(), 'Author');
assert.equal(ths[2]?.textContent.trim(), 'Genre');
assert.equal(ths[3]?.textContent.trim(), 'Publication Year');
```

`tbody` 요소 안에 하나의 `table` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table tbody'), 1);
```

`tbody` 요소는 최소 다섯 개의 행을 가져야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('tbody tr').length, 5);
```

`tbody` 요소의 각 행은 자식으로 정확히 네 개의 `td` 요소를 가져야 합니다.

```js
const rows = document.querySelectorAll('tbody tr');
assert.isAtLeast(rows.length, 1);

rows.forEach(row => {
  assert.lengthOf(row.querySelectorAll('td'), 4);
  assert.lengthOf(row.children, 4);
});
```

표 본문의 각 `td` 요소는 책 정보를 텍스트로 가져야 합니다.

```js
const tds = document.querySelectorAll('tbody tr td');
assert.isAtLeast(tds.length, 1);

tds.forEach(td => {
  assert.isAtLeast(td.textContent.trim().length, 1);
});
```

`tfoot` 요소 안에 하나의 `table` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('table tfoot'), 1);
```

`tr` 요소 안에 정확히 하나의 `tfoot` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('tfoot tr'), 1);
```

`td` 요소의 행에 있는 `tfoot` 요소는 `colspan` 속성이 `4`로 설정되어야 합니다.

```js
assert.equal(document.querySelector('tfoot tr td')?.colSpan, 4);
```

`td` 요소의 행에 있는 `tfoot` 요소는 텍스트 `Total Books: N`를 가져야 하며, `N`은 표에 있는 책 수입니다.

```js
const numberOfBooks = document.querySelectorAll('tbody tr')?.length;
assert.equal(document.querySelector('tfoot tr td').textContent.trim(), `Total Books: ${numberOfBooks}`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Book Catalog</title>
</head>

<body>
</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Book Catalog</title>
</head>
<body>
    <h1>Book Catalog</h1>
  <table>
    <thead>
      <tr>
        <th>Title</th>
        <th>Author</th>
        <th>Genre</th>
        <th>Publication Year</th>
      </tr>
    </thead>
    <tbody>
      <tr>
        <td>How to Contribute to Open-Source Projects – A Handbook for Beginners</td>
        <td>Hillary Nyakundi</td>
        <td>Open Source</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>Learn Linux for Beginners: From Basics to Advanced Techniques</td>
        <td>Zaira Hira</td>
        <td>Linux</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>How to Learn to Code and Get a Developer Job</td>
        <td>Quincy Larson</td>
        <td>Learn To Code</td>
        <td>2024</td>
      </tr>
      <tr>
        <td>The Regular Expressions Book – RegEx for JavaScript Developers</td>
        <td>Kolade Chris</td>
        <td>Regular Expressions</td>
        <td>2023</td>
      </tr>
      <tr>
        <td>The Python Code Example Handbook</td>
        <td>Farhan Hasin Chowdhury</td>
        <td>Python</td>
        <td>2023</td>
      </tr>
    </tbody>
    <tfoot>
      <tr>
        <td colspan="4">Total Books: 5</td>
      </tr>
    </tfoot>
  </table>
</body>
</html>
```
