---
id: 66be24cb4144f955b6bcc550
title: 카드 페이지 만들기
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 최소 세 장의 카드가 표시되는 웹페이지를 만들어야 합니다.
1. `main` 요소에 `playing-cards` ID가 있어야 합니다.
1. `#playing-cards` 요소 안에 클래스가 `div`인 `card` 요소가 최소 세 개 있어야 합니다.
1. 각 `.card` 요소 안에는 세 개의 `div` 요소가 있어야 하며, 첫 번째는 `left` 클래스, 두 번째는 `middle` 클래스, 세 번째는 `right` 클래스를 가져야 합니다.
1. `#playing-cards` 요소는 플렉스박스를 사용해 자식 요소들을 가로로 가운데 정렬하고, 줄 바꿈을 허용하며, 자식 요소 사이에 `20px` 간격을 두어야 합니다.
1. 각 `.card` 요소는 플렉스박스를 사용해 자식 요소들을 `space-between` 방식으로 정렬해야 합니다.
1. 각 `.left` 요소는 플렉스박스 아이템 속성을 사용해 부모 컨테이너의 시작 부분에 정렬되어야 합니다.
1. 각 `.middle` 요소는 플렉스박스 아이템 속성을 사용해 부모 컨테이너의 중앙에 정렬되어야 합니다.
1. 각 `.right` 요소는 플렉스박스 아이템 속성을 사용해 부모 컨테이너의 끝 부분에 정렬되어야 합니다.
1. 각 `.middle` 요소는 플렉스박스를 사용해 자식 요소들을 세로 열로 표시해야 합니다.

카드에 사용하고 싶다면 다음 문자를 복사해 붙여넣을 수 있습니다: `♠`, `♣`, `♥`, 그리고 `♦`.

# --hints--

`main` 요소에 `id`가 `playing-cards`인 것이 있어야 합니다.

```js
assert.exists(document.querySelector('main#playing-cards'));
```

`div` 요소 안에 클래스가 `card`인 `#playing-cards` 요소가 최소 세 개 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

각 `.card` 요소는 너비와 높이가 있어야 합니다.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

각 `.card` 요소는 정확히 세 개의 `div` 자식 요소가 있어야 합니다.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

각 `div` 요소 안에 클래스가 `left`인 `.card`가 있어야 합니다.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

각 `div` 요소 안에 클래스가 `middle`인 `.card`가 있어야 합니다.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

각 `div` 요소 안에 클래스가 `right`인 `.card`가 있어야 합니다.

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

`#playing-cards` 선택자가 요소들의 `display`를 `flex`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

`#playing-cards` 선택자가 요소들의 `justify-content`를 `center`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

`#playing-cards` 선택자가 요소들의 `flex-wrap`를 `wrap`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

`#playing-cards` 선택자가 요소들의 `gap`를 `20px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

`.card` 선택자가 요소들의 `display`를 `flex`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

`.card` 선택자가 요소들의 `justify-content`를 `space-between`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

`.left` 선택자가 요소들의 `align-self`를 `flex-start`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

`.middle` 선택자가 요소들의 `align-self`를 `center`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

`.right` 선택자가 요소들의 `align-self`를 `flex-end`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

`.middle` 선택자가 요소들의 `display`를 `flex`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

`flex-direction`가 `column`로 `.middle` 클래스에 설정되어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
