---
id: 66f686b8ebdb982fa8e14330
title: 日付変換プログラムを作成する
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

このラボでは、指定された日付を異なる形式に変換する日付変換プログラムをコードで作成します。例えば、現在の日付 `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` は以下の2つの形式でフォーマットされます。

- `(MM/DD/YYYY): 9/27/2024`。
- `(Month Day, Year): September 27, 2024`。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `currentDate` という名前の変数を作成し、`Date` オブジェクトを使って現在の日付と時刻を代入してください。
2. `currentDateFormat` という名前の変数を作成し、文字列 `Current Date and Time: [current date]` を代入してください。`[current date]` を `currentDate` 変数の結果に置き換えてください。
3. `currentDateFormat` の値をコンソールに出力してください。
4. `formatDateMMDDYYYY` という名前の関数を作成し、日付オブジェクトをパラメータとして受け取ってください。このパラメータ名は自由に決めて構いません。
5. `formatDateMMDDYYYY` 関数は現在の日付を `month/day/year` 形式に変換し、文字列 `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]` を返す必要があります。
6. `formatDateLong` という名前の関数を作成し、日付オブジェクトをパラメータとして受け取ってください。
7. `formatDateLong` 関数は現在の日付を `Month Day, Year` 形式に変換し、文字列 `Formatted Date (Month Day, Year): [formatted date]` を返す必要があります。

**注意：** テストに合格するために、日付をフォーマットするときは `en-US` をロケールとして使用してください。

# --hints--

`currentDate` オブジェクトを使って現在の日付と時刻を保持する `Date` という名前の変数を作成してください。

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

`currentDateFormat` 形式で現在の日付を保持する `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>` という名前の変数を作成してください。

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

`currentDateFormat` の値をコンソールに出力してください。

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

`formatDateMMDDYYYY` という関数を作成してください。

```js
assert.isFunction(formatDateMMDDYYYY);
```

`formatDateMMDDYYYY` 関数は1つのパラメータを受け取る必要があります。

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

日付オブジェクトが `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` を保持している場合、`formatDateMMDDYYYY` 関数は `Formatted Date (MM/DD/YYYY): 9/27/2024` を返す必要があります。

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

`formatDateLong` という関数を作成してください。

```js
assert.isFunction(formatDateLong);
```

`formatDateLong` 関数は1つのパラメータを受け取る必要があります。

```js
assert.lengthOf(formatDateLong, 1);
```

日付オブジェクトが `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` を保持している場合、`formatDateLong` 関数は `Formatted Date (Month Day, Year): September 27, 2024` を返す必要があります。

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
