---
id: 587d7fad367417b2b2512bdf
title: додати осі до візуалізації
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

Ще один спосіб покращити діаграму розсіювання — додати вісь x та вісь y.

У D3 є два методи, `axisLeft()` і `axisBottom()`, щоб відтворити вісь y та вісь x відповідно. Ось приклад створення осі x на основі `xScale` з попередніх завдань:

```js
const xAxis = d3.axisBottom(xScale);
```

Наступний крок — відтворити вісь на SVG. Для цього можна використати загальний SVG-компонент, елемент `g`. `g` означає групу. На відміну від `rect`, `circle` і `text`, вісь — це просто пряма лінія, коли її відтворюють. Оскільки це проста форма, використання `g` підходить. Останній крок — застосувати атрибут `transform`, щоб розташувати вісь у правильному місці на SVG. Інакше лінія відтворювалася б уздовж кордону SVG і була б невидимою. SVG підтримує різні типи `transforms`, але для позиціонування осі потрібен `translate`. Коли його застосовують до елемента `g`, він переміщує всю групу вправо і вниз на задані значення. Ось приклад:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

Вищенаведений код розміщує вісь x внизу SVG. Потім її передають як аргумент методу `call()`. Вісь y працює так само, за винятком того, що аргумент `translate` має форму `(x, 0)`. Оскільки `translate` є рядком у методі `attr()` вище, ви можете використовувати конкатенацію, щоб включити змінні значення для його аргументів.

# --instructions--

Діаграма розсіювання тепер має вісь x. Створіть вісь y у змінній з ім’ям `yAxis` за допомогою методу `axisLeft()`. Потім відтворіть вісь, використовуючи елемент `g`. Переконайтеся, що ви застосували атрибут `transform`, щоб зсунути вісь вправо на кількість одиниць відступу, і `0` одиниць вниз. Не забудьте `call()` відтворити вісь.

# --hints--

Ваш код має використовувати метод `axisLeft()` з аргументом `yScale`.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

Елемент `g` для осі y має мати атрибут `transform`, щоб зсунути вісь на `(60, 0)`.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

Ваш код має викликати `yAxis`.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
