---
id: 587d7fa9367417b2b2512bce
title: Динамічно встановити координати для кожної гістограми
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

Останнє завдання створило та додало прямокутник до елемента `svg` для кожної точки в `dataset`, щоб представити гістограму. На жаль, вони всі були накладені один на одного.

Розміщення прямокутника контролюється атрибутами `x` та `y`. Вони вказують D3, де починати малювати форму в області `svg`. В останньому завданні вони були встановлені в 0, тому кожна гістограма розміщувалась у верхньому лівому куті.

Для гістограми всі стовпчики повинні стояти на одному вертикальному рівні, що означає, що значення `y` залишається однаковим (на 0) для всіх стовпчиків. Однак значення `x` має змінюватися при додаванні нових стовпчиків. Пам’ятайте, що більші значення `x` відсувають елементи далі вправо. Під час ітерування над елементами масиву в `dataset` значення `x` має збільшуватися.

Метод `attr()` у D3 приймає функцію зворотного виклику, щоб динамічно встановити цей атрибут. Функція зворотного виклику приймає два аргументи: один для самої точки даних (зазвичай `d`) і один для індексу точки даних у масиві. Другий аргумент для індексу є необов’язковим. Ось формат:

```js
selection.attr('property', (d, i) => {});
```

Важливо зазначити, що вам НЕ потрібно писати цикл `for` або використовувати `forEach()` для ітерування над елементами набору даних. Пам’ятайте, що метод `data()` аналізує набір даних, і будь-який метод, який виконується після `data()`, запускається один раз для кожного елемента набору даних.

# --instructions--

Змініть функцію зворотного виклику атрибута `x`, щоб вона повертала індекс, помножений на 30.

**Примітка:** Кожен стовпчик має ширину 25, тому збільшення кожного значення `x` на 30 додає трохи простору між стовпчиками. Будь-яке значення більше за 25 підійде в цьому прикладі.

# --hints--

Перший `rect` має мати значення `x` рівне `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

Другий `rect` має мати значення `x` рівне `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

Третій `rect` має мати значення `x` рівне `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

Четвертий `rect` має мати значення `x` рівне `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

П’ятий `rect` має мати значення `x` рівне `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

Шостий `rect` має мати значення `x` рівне `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

Сьомий `rect` має мати значення `x` рівне `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

Восьмий `rect` має мати значення `x` рівне `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

Дев’ятий `rect` має мати значення `x` рівне `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
