---
id: 587d7fa6367417b2b2512bc3
title: Вибрати групу елементів за допомогою D3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

У D3 також є метод `selectAll()` для вибору групи елементів. Він повертає масив HTML-вузлів для всіх елементів у документі, які відповідають вхідному рядку. Ось приклад, як вибрати всі теги посилань у документі:

```js
const anchors = d3.selectAll('a');
```

Як і метод `select()`, `selectAll()` підтримує ланцюжкове викликання методів, і ви можете використовувати його з іншими методами.

# --instructions--

Виберіть усі теги `li` у документі та змініть їхній текст на рядок `list item`, використовуючи ланцюжок викликів методу `.text()`.

# --hints--

На сторінці має бути 3 елементи `li`, і текст у кожному з них має бути `list item`. Великі літери та пропуски мають точно співпадати.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Ваш код має отримувати доступ до об’єкта `d3`.

```js
assert.match(code, /d3/g);
```

Ваш код має використовувати метод `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
