---
id: 587d7fa7367417b2b2512bc5
title: Робота з динамічними даними в D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

Останні два завдання охоплюють основи динамічного відображення даних за допомогою D3, використовуючи методи `data()` та `enter()`. Ці методи приймають множину даних і разом із методом `append()` створюють новий DOM-елемент для кожного запису в множині даних.

У попередньому завданні ви створили новий `h2` елемент для кожного елемента в масиві `dataset`, але всі вони містили однаковий текст `New Title`. Це тому, що ви не використали дані, які прив’язані до кожного з `h2` елементів.

Метод `text()` у D3 може приймати рядок або функцію зворотного виклику як аргумент:

```js
selection.text(d => d);
```

У наведеному вище прикладі параметр `d` посилається на один запис у множині даних, до якої прив’язаний вибір.

Використовуючи поточний приклад як контекст, перший `h2` елемент прив’язаний до 12, другий `h2` елемент прив’язаний до 31, третій `h2` елемент прив’язаний до 22 і так далі.

# --instructions--

Змініть метод `text()` так, щоб кожен `h2` елемент відображав відповідне значення з масиву `dataset` з одним пробілом і рядком `USD`. Наприклад, перший заголовок має бути `12 USD`.

# --hints--

Перший `h2` має містити текст `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

Другий `h2` має містити текст `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

Третій `h2` має містити текст `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

Четвертий `h2` має містити текст `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

П’ятий `h2` має містити текст `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

Шостий `h2` має містити текст `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

Сьомий `h2` має містити текст `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

Восьмий `h2` має містити текст `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

Дев’ятий `h2` має містити текст `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
