---
id: 587d7fac367417b2b2512bdd
title: Використання динамічних масштабів
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

Методи D3 `min()` і `max()` корисні для налаштування масштабу.

При роботі зі складним набором даних одним із пріоритетів є встановлення масштабу так, щоб візуалізація поміщалася у ширину та висоту контейнера SVG. Ви хочете, щоб усі дані були побудовані на графіку всередині SVG і були видимі на вебсторінці.

Приклад нижче встановлює масштаб осі x для даних точкової діаграми. Метод `domain()` передає інформацію масштабу про сирі значення даних для побудови. Метод `range()` дає інформацію про фактичний простір на вебсторінці для візуалізації.

У прикладі домен йде від 0 до максимального значення в множині. Використовується метод `max()` з функцією зворотного виклику, що базується на значеннях x у масивах. Діапазон використовує ширину SVG (`w`), але також включає відступ. Це створює простір між точками точкової діаграми та краєм SVG.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

Відступ може спочатку збивати з пантелику. Уявіть вісь x як горизонтальну лінію від 0 до 500 (значення ширини для SVG). Включення відступу в метод `range()` змушує побудову починатися з 30 по цій лінії (замість 0) і закінчуватися на 470 (замість 500).

# --instructions--

Використайте змінну `yScale` для створення лінійного масштабу осі y. Домен має починатися з нуля і йти до максимального значення `y` у множині. Діапазон має використовувати висоту SVG (`h`) і включати відступ.

**Примітка:** Пам’ятайте, що побудова має бути правильною за напрямком. Коли ви встановлюєте діапазон для координат y, більшим аргументом (висота мінус відступ) є перший аргумент, а меншим — другий.

# --hints--

Текст у `h2` має бути `30`.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

`domain()` для yScale має бути еквівалентним `[0, 411]`.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

`range()` для yScale має бути еквівалентним `[470, 30]`.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
