---
id: 587d7fa7367417b2b2512bc7
title: змінювати стилі залежно від даних
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3 — це про візуалізацію та подання даних. Ймовірно, ви захочете змінювати стилі елементів залежно від даних.
Наприклад, ви можете захотіти пофарбувати точку даних у синій, якщо її значення менше за 20, і в червоний інакше. Ви можете використати функцію зворотного виклику в методі `style()` і включити умовну логіку. Функція зворотного виклику використовує параметр `d` для представлення точки даних:

```js
selection.style('color', d => {});
```

Метод `style()` не обмежується встановленням `color` — його можна використовувати й з іншими CSS-властивостями.

# --instructions--

Додайте метод `style()` до коду в редакторі, щоб умовно встановити `color` елементів `h2`. Напишіть функцію зворотного виклику так, щоб якщо значення даних менше за 20, вона повертала червоний, інакше — зелений.

**Примітка:** Ви можете використати логіку if-else або тернарний оператор.

# --hints--

Перший `h2` має мати `color` червоного кольору.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Другий `h2` має мати `color` зеленого кольору.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Третій `h2` має мати `color` зеленого кольору.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Четвертий `h2` має мати `color` червоного кольору.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

П’ятий `h2` має мати `color` зеленого кольору.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Шостий `h2` має мати `color` червоного кольору.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Сьомий `h2` має мати `color` зеленого кольору.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Восьмий `h2` має мати `color` червоного кольору.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Дев’ятий `h2` має мати `color` червоного кольору.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
