---
id: 587d7fab367417b2b2512bd9
title: додати мітки до точок на точковій діаграмі
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

Ви можете додати текст, щоб створити мітки для точок на точковій діаграмі.

Мета — відобразити значення, розділені комою, для першого (`x`) і другого (`y`) полів кожного елемента в `dataset`.

Вузли `text` потребують атрибутів `x` і `y` для позиціювання на SVG. У цьому завданні значення `y` (яке визначає висоту) може використовувати те саме значення, що й `circle` для свого атрибута `cy`. Значення `x` може бути трохи більшим за значення `cx` у `circle`, щоб мітка була видимою. Це зсунеться мітку праворуч від побудованої на графіку точки.

# --instructions--

Позначте кожну точку на точковій діаграмі, використовуючи елементи `text`. Текст мітки має містити два значення, розділені комою та пробілом. Наприклад, мітка для першої точки — `34, 78`. Встановіть атрибут `x` так, щоб він був на `5` одиниць більшим за значення, яке ви використали для атрибута `cx` у `circle`. Встановіть атрибут `y` так само, як використовується для значення `cy` у `circle`.

# --hints--

Ваш код має містити 10 елементів `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Перша мітка має текст `34, 78`, значення `x` дорівнює `39`, а значення `y` — `422`.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

Друга мітка має текст `109, 280`, значення `x` дорівнює `114`, а значення `y` — `220`.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

Третя мітка має текст `310, 120`, значення `x` дорівнює `315`, а значення `y` — `380`.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

Четверта мітка має текст `79, 411`, значення `x` дорівнює `84`, а значення `y` — `89`.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

П’ята мітка має текст `420, 220`, значення `x` дорівнює `425`, а значення `y` — `280`.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

Шоста мітка має текст `233, 145`, значення `x` дорівнює `238`, а значення `y` — `355`.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

Сьома мітка має текст `333, 96`, значення `x` дорівнює `338`, а значення `y` — `404`.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

Восьма мітка має текст `222, 333`, значення `x` дорівнює `227`, а значення `y` — `167`.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

Дев’ята мітка має текст `78, 320`, значення `x` дорівнює `83`, а значення `y` — `180`.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

Десята мітка має текст `21, 123`, значення `x` дорівнює `26`, а значення `y` — `377`.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
