---
id: 587d7fa9367417b2b2512bd0
title: Інвертувати елементи SVG
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

Ви могли помітити, що гістограма виглядає ніби перевернута або інвертована. Це через те, як SVG використовує координати (x, y).

В SVG початкова точка координат знаходиться у верхньому лівому куті. Координата `x` зі значенням 0 розміщує форму на лівому краю області SVG. Координата `y` зі значенням 0 розміщує форму на верхньому краю області SVG. Вищі значення `x` зсувають прямокутник вправо. Вищі значення `y` зсувають прямокутник вниз.

Щоб зробити смужки правильно орієнтованими, потрібно змінити спосіб обчислення координати `y`. Вона має враховувати як висоту смужки, так і загальну висоту області SVG.

Висота області SVG дорівнює 100. Якщо у множині даних є точка зі значенням 0, ви хочете, щоб смужка починалася знизу області SVG (а не зверху). Для цього координата `y` має мати значення 100. Якщо значення точки даних дорівнює 1, ви починаєте з координати `y` зі значенням 100, щоб розмістити смужку внизу. Потім потрібно врахувати висоту смужки 1, тому фінальна координата `y` буде 99.

Координата `y`, яка є `y = heightOfSVG - heightOfBar`, розмістить смужки правильно орієнтованими.

# --instructions--

Змініть функцію зворотного виклику для атрибута `y`, щоб розмістити смужки правильно орієнтованими. Пам’ятайте, що `height` смужки дорівнює 3-кратному значенню даних `d`.

**Примітка:** Загалом, залежність має вигляд `y = h - m * d`, де `m` — це константа, що масштабує точки даних.

# --hints--

Перша `rect` має мати значення `y` рівне `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

Друга `rect` має мати значення `y` рівне `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

Третя `rect` має мати значення `y` рівне `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

Четверта `rect` має мати значення `y` рівне `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

П’ята `rect` має мати значення `y` рівне `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

Шоста `rect` має мати значення `y` рівне `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

Сьома `rect` має мати значення `y` рівне `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

Восьма `rect` має мати значення `y` рівне `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

Дев’ята `rect` має мати значення `y` рівне `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
