---
id: 587d7fa8367417b2b2512bca
title: Змінити подання гістограми
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

Останнє завдання створило гістограму, але є кілька змін форматування, які можуть її покращити:

1. Додати пробіл між кожною смугою, щоб візуально їх розділити, що робиться шляхом додавання поля в CSS для класу `bar`

2. Збільшити висоту смуг, щоб краще показати різницю в значеннях, що робиться шляхом помноження значення на число для масштабування висоти

# --instructions--

Спочатку додайте `margin` зі значенням `2px` до класу `bar` у тегу `style`. Потім змініть функцію зворотного виклику в методі `style()`, щоб вона повертала значення, помножене `10` на початкове значення даних (плюс `px`).

**Примітка:** Помноження кожної точки даних на _одну й ту ж_ константу лише змінює масштаб. Це як збільшення зображення, і не змінює значення основних даних.

# --hints--

Перший `div` повинен мати поле `height` розміром `120` пікселів і `margin` розміром `2` пікселів.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

Другий `div` повинен мати поле `height` розміром `310` пікселів і `margin` розміром `2` пікселів.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

Третій `div` повинен мати поле `height` розміром `220` пікселів і `margin` розміром `2` пікселів.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

Четвертий `div` повинен мати поле `height` розміром `170` пікселів і `margin` розміром `2` пікселів.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

П’ятий `div` повинен мати поле `height` розміром `250` пікселів і `margin` розміром `2` пікселів.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

Шостий `div` повинен мати поле `height` розміром `180` пікселів і `margin` розміром `2` пікселів.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

Сьомий `div` повинен мати поле `height` розміром `290` пікселів і `margin` розміром `2` пікселів.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

Восьмий `div` повинен мати поле `height` розміром `140` пікселів і `margin` розміром `2` пікселів.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

Дев’ятий `div` повинен мати поле `height` розміром `90` пікселів і `margin` розміром `2` пікселів.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
