---
id: 587d7fac367417b2b2512bdb
title: Встановлення домену та діапазону на шкалі
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

За замовчуванням шкали використовують ідентичне відношення. Це означає, що вхідне значення спроєктоване у вихідне значення. Однак шкали можуть бути набагато гнучкішими та цікавішими.

Припустімо, що в наборі даних є значення в діапазоні від 50 до 480. Це вхідна інформація для шкали, також відома як <dfn>домен</dfn>.

Ви хочете спроєктувати ці точки вздовж осі `x` на SVG, між 10 одиницями та 500 одиницями. Це вихідна інформація, також відома як <dfn>діапазон</dfn>.

Методи `domain()` і `range()` встановлюють ці значення для шкали. Обидва методи приймають масив щонайменше з двох елементів як аргумент. Ось приклад:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

Послідовно у консолі будуть виведені такі значення: `10`, `500`, `323.37` і `807.67`.

Зверніть увагу, що шкала використовує лінійне відношення між значеннями домену та діапазону, щоб визначити, яким має бути вихідне значення для заданого числа. Мінімальне значення в домені (50) спроєктоване у мінімальне значення (10) в діапазоні.

# --instructions--

Створіть шкалу та встановіть її домен у `[250, 500]`, а діапазон у `[10, 150]`.

**Примітка:** Ви можете ланцюжком викликати методи `domain()` і `range()` на змінній `scale`.

# --hints--

Ваш код має використовувати метод `domain()`.

```js
assert.match(code, /\.domain/g);
```

`domain()` `scale` має бути встановлено в `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Ваш код має використовувати метод `range()`.

```js
assert.match(code, /\.range/g);
```

`range()` `scale` має бути встановлено в `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Текст у `h2` має бути `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
