---
id: 587d7faa367417b2b2512bd2
title: Додайте мітки до елементів D3
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

D3 дозволяє додавати мітки до елемента графіка, наприклад, до стовпчика, використовуючи SVG-елемент `text`.

Як і елемент `rect`, елемент `text` повинен мати атрибути `x` і `y`, щоб розмістити його на SVG. Також він повинен мати доступ до даних, щоб відобразити ці значення.

D3 дає вам високий рівень контролю над тим, як ви позначаєте свої стовпчики.

# --instructions--

Код у редакторі вже зв’язує дані з кожним новим елементом `text`. Спочатку додайте вузли `text` до `svg`. Потім додайте атрибути для координат `x` і `y`. Вони мають обчислюватися так само, як і для `rect`, за винятком того, що значення `y` для `text` має розмістити мітку на 3 одиниці вище за стовпчик. Нарешті, використайте метод D3 `text()`, щоб встановити мітку рівною значенню точки даних.

**Примітка:** Щоб мітка була вище за стовпчик, визначте, чи має значення `y` для `text` бути на 3 більше або на 3 менше за значення `y` для стовпчика.

# --hints--

Перший елемент `text` має мати мітку `12` і значення `y` рівне `61`.

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

Другий елемент `text` має мати мітку `31` і значення `y` рівне `4`.

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

Третій елемент `text` має мати мітку `22` і значення `y` рівне `31`.

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

Четвертий елемент `text` має мати мітку `17` і значення `y` рівне `46`.

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

П’ятий елемент `text` має мати мітку `25` і значення `y` рівне `22`.

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

Шостий елемент `text` має мати мітку `18` і значення `y` рівне `43`.

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

Сьомий елемент `text` має мати мітку `29` і значення `y` рівне `10`.

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

Восьмий елемент `text` має мати мітку `14` і значення `y` рівне `55`.

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

Дев’ятий елемент `text` має мати мітку `9` і значення `y` рівне `70`.

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```
