---
id: 587d7fa8367417b2b2512bc9
title: Оновити висоту елемента динамічно
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

Попередні завдання охоплювали, як відображати дані з масиву та як додавати CSS-класи. Ви можете поєднати ці уроки, щоб створити просту гістограму. Це складається з двох кроків:

1. Створити `div` для кожної точки даних у масиві

2. Надати кожному `div` динамічну висоту, використовуючи функцію зворотного виклику в методі `style()`, яка встановлює висоту рівною значенню даних

Нагадаємо формат встановлення стилю за допомогою функції зворотного виклику:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

Додайте метод `style()` до коду в редакторі, щоб встановити властивість `height` для кожного елемента. Використовуйте функцію зворотного виклику, щоб повернути значення точки даних з доданим рядком `px`.

# --hints--

Перший `div` має мати властивість `height` зі значенням `12` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

Другий `div` має мати властивість `height` зі значенням `31` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

Третій `div` має мати властивість `height` зі значенням `22` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

Четвертий `div` має мати властивість `height` зі значенням `17` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

П’ятий `div` має мати властивість `height` зі значенням `25` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

Шостий `div` має мати властивість `height` зі значенням `18` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

Сьомий `div` має мати властивість `height` зі значенням `29` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

Восьмий `div` має мати властивість `height` зі значенням `14` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

Дев’ятий `div` має мати властивість `height` зі значенням `9` пікселів.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
