---
id: 587d7fa8367417b2b2512bcd
title: створити смужку для кожної точки даних у множині
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

У попередньому завданні було додано лише один прямокутник до елемента `svg`, щоб представити смужку. Тут ви поєднаєте те, що дізналися досі про `data()`, `enter()` і SVG-форми, щоб створити та додати прямокутник для кожної точки даних у `dataset`.

У попередньому завданні показано формат створення та додавання `div` для кожного елемента в `dataset`:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

Існують деякі відмінності при роботі з елементами `rect` замість `div`. Елементи `rect` потрібно додавати до елемента `svg`, а не безпосередньо до `body`. Також потрібно вказати D3, де розмістити кожен `rect` у області `svg`. Розміщення смужок буде розглянуто в наступному завданні.

# --instructions--

Використайте методи `data()`, `enter()` і `append()`, щоб створити та додати `rect` для кожного елемента в `dataset`. Смужки мають відображатися всі одна на одній; це буде виправлено в наступному завданні.

# --hints--

У вашому документі має бути 9 елементів `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

Ваш код має використовувати метод `data()`.

```js
assert.match(code, /\.data/g);
```

Ваш код має використовувати метод `enter()`.

```js
assert.match(code, /\.enter/g);
```

Ваш код має використовувати метод `append()`.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
