---
id: 587d7fa8367417b2b2512bcb
title: Дізнайтеся про SVG у D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> означає <dfn>масштабована векторна графіка</dfn>.

Тут «масштабована» означає, що якщо ви збільшуєте або зменшуєте об’єкт, він не буде пікселізованим. Він масштабується відповідно до системи відображення, чи то на маленькому екрані мобільного пристрою, чи на великому телевізійному моніторі.

SVG використовується для створення загальних геометричних форм. Оскільки D3 спроєктований для представлення даних у візуальній формі, він використовує SVG для створення форм для візуалізації. SVG-форми для вебсторінки мають розміщуватися всередині HTML-тегу `svg`.

CSS може бути масштабованим, коли стилі використовують відносні одиниці (наприклад, `vh`, `vw` або відсотки), але використання SVG є гнучкішим для створення візуалізації даних.

# --instructions--

Додайте вузол `svg` до `body` за допомогою `append()`. Додайте йому атрибут `width`, встановлений у надану константу `w`, і атрибут `height`, встановлений у надану константу `h` за допомогою методів `attr()` або `style()` для кожного. Ви побачите це у вихідних даних, оскільки до нього застосовано рожевий `background-color` у тегу `style`.

**Примітка:** При використанні атрибутів ширини та висоти `attr()` не мають одиниць виміру. Це основа масштабування — елемент завжди матиме співвідношення ширини до висоти 5:1, незалежно від рівня збільшення.

# --hints--

Ваш документ має містити 1 елемент `svg`.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

Елемент `svg` має мати атрибут `width`, встановлений у `500`, або бути стилізованим із шириною `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

Елемент `svg` має мати атрибут `height`, встановлений у `100`, або бути стилізованим із висотою `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
