---
id: 587d7fac367417b2b2512bde
title: Використання заздалегідь визначеної шкали для розміщення елементів
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

Після налаштування шкал настав час знову спроєктувати діаграму розсіювання. Шкали — це ніби функції обробки, які перетворюють сирі дані `x` і `y` у значення, що підходять і правильно відтворюються на SVG. Вони тримають дані в межах області побудови на екрані.

Ви встановлюєте значення атрибутів координат для SVG-форми за допомогою функції масштабування. Це включає атрибути `x` і `y` для елементів `rect` або `text`, або `cx` і `cy` для `circles`. Ось приклад:

```js
shape.attr('x', d => xScale(d[0]));
```

Шкали встановлюють атрибути координат форми, щоб розмістити точки даних на SVG. Вам не потрібно застосовувати шкали, коли ви відображаєте фактичне значення даних, наприклад, у методі `text()` для підказки або мітки.

# --instructions--

Використайте `xScale` і `yScale`, щоб розмістити як форми `circle`, так і `text` на SVG. Для `circles` застосуйте шкали, щоб встановити атрибути `cx` і `cy`. Також задайте їм радіус `5` одиниць.

Для елементів `text` застосуйте шкали, щоб встановити атрибути `x` і `y`. Мітки мають бути зміщені праворуч від точок. Для цього додайте `10` одиниць до значення даних `x` перед передачею його в `xScale`.

# --hints--

Ваш код має містити 10 елементів `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Перший елемент `circle` має мати значення `cx` приблизно `91` і значення `cy` приблизно `368` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

Другий елемент `circle` має мати значення `cx` приблизно `159` і значення `cy` приблизно `181` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

Третій елемент `circle` має мати значення `cx` приблизно `340` і значення `cy` приблизно `329` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

Четвертий елемент `circle` має мати значення `cx` приблизно `131` і значення `cy` приблизно `60` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

П’ятий елемент `circle` має мати значення `cx` приблизно `440` і значення `cy` приблизно `237` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

Шостий елемент `circle` має мати значення `cx` приблизно `271` і значення `cy` приблизно `306` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

Сьомий елемент `circle` має мати значення `cx` приблизно `361` і значення `cy` приблизно `351` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

Восьмий елемент `circle` має мати значення `cx` приблизно `261` і значення `cy` приблизно `132` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

Дев’ятий елемент `circle` має мати значення `cx` приблизно `131` і значення `cy` приблизно `144` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

Десятий елемент `circle` має мати значення `cx` приблизно `79` і значення `cy` приблизно `326` після застосування шкал. Також він має мати значення `r` рівне `5`.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

Ваш код має містити 10 елементів `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Перша мітка має мати значення `x` приблизно `100` і значення `y` приблизно `368` після застосування шкал.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

Друга мітка має мати значення `x` приблизно `168` і значення `y` приблизно `181` після застосування шкал.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

Третя мітка має мати значення `x` приблизно `350` і значення `y` приблизно `329` після застосування шкал.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

Четверта мітка має мати значення `x` приблизно `141` і значення `y` приблизно `60` після застосування шкал.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

П’ята мітка має мати значення `x` приблизно `449` і значення `y` приблизно `237` після застосування шкал.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

Шоста мітка має мати значення `x` приблизно `280` і значення `y` приблизно `306` після застосування шкал.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

Сьома мітка має мати значення `x` приблизно `370` і значення `y` приблизно `351` після застосування шкал.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

Восьма мітка має мати значення `x` приблизно `270` і значення `y` приблизно `132` після застосування шкал.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

Дев’ята мітка має мати значення `x` приблизно `140` і значення `y` приблизно `144` після застосування шкал.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

Десята мітка має мати значення `x` приблизно `88` і значення `y` приблизно `326` після застосування шкал.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
