---
id: 587d7fab367417b2b2512bd7
title: Створіть точкову діаграму за допомогою кіл SVG
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

Точкова діаграма — це ще один тип візуалізації. Зазвичай вона використовує кола для спроєктування точок даних, кожна з яких має два значення. Ці значення пов’язані з осями `x` та `y` і використовуються для позиціювання кола у візуалізації.

SVG має тег `circle` для створення форми кола. Він працює дуже схоже на елементи `rect`, які ви використовували для гістограми.

# --instructions--

Використайте методи `data()`, `enter()` та `append()` для прив’язування `dataset` до нових елементів `circle`, які додані до SVG.

**Примітка:** Кола не будуть видимі, бо ми ще не встановили їхні атрибути. Це зробимо в наступному завданні.

# --hints--

У вашому коді має бути 10 елементів `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
