---
id: 587d7fa9367417b2b2512bcf
title: Динамічно змінюйте висоту кожної смужки
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

Висоту кожної смужки можна встановити відповідно до значення точки даних у масиві, подібно до того, як значення `x` було встановлено динамічно.

```js
selection.attr('property', (d, i) => {});
```

Тут `d` буде значенням точки даних, а `i` — індексом точки даних у масиві.

# --instructions--

Змініть функцію зворотного виклику для атрибута `height`, щоб вона повертала значення даних, помножене на 3.

**Примітка:** Пам’ятайте, що множення всіх точок даних на одну й ту ж константу масштабує дані (як збільшення масштабу). Це допомагає краще бачити відмінності між значеннями смужок у цьому прикладі.

# --hints--

Перший `rect` має мати атрибут `height` зі значенням `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

Другий `rect` має мати атрибут `height` зі значенням `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

Третій `rect` має мати атрибут `height` зі значенням `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

Четвертий `rect` має мати атрибут `height` зі значенням `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

П’ятий `rect` має мати атрибут `height` зі значенням `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

Шостий `rect` має мати атрибут `height` зі значенням `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

Сьомий `rect` має мати атрибут `height` зі значенням `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

Восьмий `rect` має мати атрибут `height` зі значенням `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

Дев’ятий `rect` має мати атрибут `height` зі значенням `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
