---
id: 587d7fab367417b2b2512bd8
title: Додайте атрибути до елементів кола
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

Останнє завдання створило `circle` елементи для кожної точки в `dataset` і додало їх до SVG. Але D3 потрібна додаткова інформація про позицію та розмір кожного `circle`, щоб відобразити їх правильно.

`circle` в SVG має три основні атрибути. Атрибути `cx` і `cy` — це координати. Вони вказують D3, де розташувати _центр_ форми на SVG. Радіус (атрибут `r`) задає розмір `circle`.

Як і координата `rect` `y`, атрибут `cy` для `circle` вимірюється від верхньої частини SVG, а не від нижньої.

Усі три атрибути можуть використовувати функцію зворотного виклику, щоб динамічно встановлювати свої значення. Пам’ятайте, що всі методи, які викликаються після `data(dataset)`, запускаються один раз для кожного елемента в `dataset`. Параметр `d` у функції зворотного виклику посилається на поточний елемент у `dataset`, який є масивом для кожної точки. Ви використовуєте дужкову нотацію, як `d[0]`, щоб отримати доступ до значень у цьому масиві.

# --instructions--

Додайте атрибути `cx`, `cy` і `r` до елементів `circle`. Значення `cx` має бути першим числом у масиві для кожного елемента в `dataset`. Значення `cy` має базуватися на другому числі в масиві, але переконайтеся, що діаграма відображається правильно, а не перевернутою. Значення `r` має бути `5` для всіх кіл.

# --hints--

У вашому коді має бути 10 елементів `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Перший елемент `circle` має мати значення `cx` рівне `34`, значення `cy` рівне `422` і значення `r` рівне `5`.

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

Другий елемент `circle` має мати значення `cx` рівне `109`, значення `cy` рівне `220` і значення `r` рівне `5`.

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

Третій елемент `circle` має мати значення `cx` рівне `310`, значення `cy` рівне `380` і значення `r` рівне `5`.

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

Четвертий елемент `circle` має мати значення `cx` рівне `79`, значення `cy` рівне `89` і значення `r` рівне `5`.

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

П’ятий елемент `circle` має мати значення `cx` рівне `420`, значення `cy` рівне `280` і значення `r` рівне `5`.

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

Шостий елемент `circle` має мати значення `cx` рівне `233`, значення `cy` рівне `355` і значення `r` рівне `5`.

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

Сьомий елемент `circle` має мати значення `cx` рівне `333`, значення `cy` рівне `404` і значення `r` рівне `5`.

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

Восьмий елемент `circle` має мати значення `cx` рівне `222`, значення `cy` рівне `167` і значення `r` рівне `5`.

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

Дев’ятий елемент `circle` має мати значення `cx` рівне `78`, значення `cy` рівне `180` і значення `r` рівне `5`.

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

Десятий елемент `circle` має мати значення `cx` рівне `21`, значення `cy` рівне `377` і значення `r` рівне `5`.

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
