---
id: 587d7fac367417b2b2512bdc
title: >-
  Використання функцій d3.max і d3.min для знаходження мінімальних і максимальних значень у наборі даних
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

Методи D3 `domain()` і `range()` встановлюють цю інформацію для вашої шкали на основі даних. Існує кілька методів, які полегшують це.

Часто, коли ви встановлюєте домен, ви хочете використати мінімальні та максимальні значення в наборі даних. Спроба знайти ці значення вручну, особливо у великому наборі даних, може призвести до помилок.

У D3 є два методи — `min()` і `max()`, які повертають цю інформацію. Ось приклад:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

Набір даних може містити вкладені масиви, як `[x, y]` координатні пари, які були у прикладі з точковою діаграмою. У такому випадку потрібно вказати D3, як обчислювати максимум і мінімум. На щастя, обидва методи `min()` і `max()` приймають функцію зворотного виклику. У цьому прикладі аргумент функції зворотного виклику `d` відповідає поточному внутрішньому масиву. Функція зворотного виклику має повернути елемент із внутрішнього масиву (значення `x` або `y`), за яким ви хочете обчислити максимум або мінімум. Ось приклад, як знайти мінімальні та максимальні значення у масиві масивів:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` матиме значення `1`.

# --instructions--

Масив `positionData` містить підмасиви координат x, y і z. Використайте метод D3, щоб знайти максимальне значення координати z (третє значення) з масивів і збережіть його у змінній `output`.

# --hints--

Текст у `h2` має бути `8`.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

Ваш код має використовувати метод `max()`.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
