---
id: 587d7fa8367417b2b2512bcc
title: Відображення форм за допомогою SVG
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

Останнє завдання створило `svg` елемент із заданою шириною та висотою, який був видимим, бо до нього було застосовано `background-color` у `style` тегу. Код виділив місце для заданої ширини та висоти.

Наступний крок — створити форму, щоб розмістити її в `svg` області. Існує кілька підтримуваних форм у SVG, таких як прямокутники та кола. Їх використовують для відображення даних. Наприклад, прямокутник (`<rect>`) у SVG може створити стовпець на гістограмі.

Коли ви розміщуєте форму в `svg` області, можна вказати її положення за допомогою координат `x` та `y`. Початкова точка (0, 0) знаходиться у верхньому лівому куті. Позитивні значення для `x` зсувають форму вправо, а позитивні значення для `y` — вниз від початкової точки.

Щоб розмістити форму посередині `svg` розміром 500 (ширина) x 100 (висота) з останнього завдання, координата `x` буде 250, а координата `y` — 50.

SVG `rect` має чотири атрибути. Це координати `x` та `y`, які вказують, де вона розміщена в `svg` області. Також є `height` та `width`, щоб задати розмір.

# --instructions--

Додайте `rect` форму до `svg` за допомогою `append()`, і надайте їй атрибут `width` зі значенням `25` та атрибут `height` зі значенням `100`. Також задайте `rect` атрибути `x` і `y`, кожен зі значенням `0`.

# --hints--

Ваш документ повинен містити 1 `rect` елемент.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

Елемент `rect` повинен мати атрибут `width`, встановлений у `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

Елемент `rect` повинен мати атрибут `height`, встановлений у `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

Елемент `rect` повинен мати атрибут `x`, встановлений у `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

Елемент `rect` повинен мати атрибут `y`, встановлений у `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
