---
id: 587d7faa367417b2b2512bd6
title: додати підказку до елемента D3
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

Підказка показує більше інформації про елемент на сторінці, коли користувач наводить курсор на цей елемент. Існує кілька способів додати підказку до візуалізації. У цьому завданні використовується SVG `title` елемент.

`title` поєднується з методом `text()` для динамічного додавання даних до стовпчиків.

# --instructions--

Додайте `title` елемент під кожним `rect` вузлом. Потім викличте метод `text()` з функцією зворотного виклику, щоб текст відображав значення даних.

# --hints--

У вашому коді має бути 9 `title` елементів.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

Перший `title` елемент має мати текст підказки `12`.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

Другий `title` елемент має мати текст підказки `31`.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

Третій `title` елемент має мати текст підказки `22`.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

Четвертий `title` елемент має мати текст підказки `17`.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

П’ятий `title` елемент має мати текст підказки `25`.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

Шостий `title` елемент має мати текст підказки `18`.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

Сьомий `title` елемент має мати текст підказки `29`.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

Восьмий `title` елемент має мати текст підказки `14`.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

Дев’ятий `title` елемент має мати текст підказки `9`.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
