---
id: 587d7fa7367417b2b2512bc8
title: додати класи за допомогою D3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

Використовувати багато рядкових стилів на HTML-елементах стає важко керувати, навіть у невеликих застосунках. Легше додати клас до елементів і стилізувати цей клас один раз за допомогою CSS-правил. У D3 є метод `attr()`, щоб додати будь-який HTML-атрибут до елемента, включно з іменем класу.

Метод `attr()` працює так само, як і `style()`. Він приймає значення, розділені комами, і може використовувати функцію зворотного виклику. Ось приклад, як додати клас `container` до вибірки:

```js
selection.attr('class', 'container');
```

Зверніть увагу, що параметр `class` залишатиметься однаковим щоразу, коли потрібно додати клас, і лише параметр `container` змінюватиметься.

# --instructions--

Додайте метод `attr()` до коду в редакторі та додайте клас `bar` до елементів `div`.

# --hints--

Ваші елементи `div` повинні мати клас `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Ваш код має використовувати метод `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
