---
id: 66faac4139dbbd5dd632c860
title: Створіть форму заявки на працевлаштування
challengeType: 25
dashedName: lab-job-application-form
demoType: onClick
---

# --description--

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `div` з класом `container`.
1. Елемент `div` має містити елемент `form`.
1. Форма має містити елемент `input` з типом `text` та `id` зі значенням `name`, де користувач буде вводити повне ім’я.
1. Ви повинні мати ще один елемент `input` з типом `email` та `id` зі значенням `email`, де користувач буде вводити електронну адресу.
1. Форма має містити елемент `select` з `id` зі значенням `position`, де користувач зможе обрати посаду.
1. Ви повинні мати елемент `fieldset` з класом `radio-group`.
1. В межах `.radio-group` має бути група елементів `input` з типом `radio` та відповідними мітками для вибору зайнятості (наприклад, Full-Time і Part-Time). Назвою цієї групи має бути `availability`.
1. Ви повинні мати елемент `textarea` з `id` зі значенням `message`, де користувач буде вводити повідомлення.
1. You should associate every `input` element with a `label` element.
1. Ви повинні мати елемент `button` з типом `submit`, щоб користувач надіслав форму.
1. Додайте псевдоклас `:focus` до елементів `input` та `textarea`, щоб колір кордону змінювався під час фокусування.
1. Елементи `input`, `select` та `textarea` повинні мати псевдоклас `:invalid`, який змінює колір кордону на червоний, якщо введено недійсне значення.
1. Елементи `input`, `select` та `textarea` повинні мати псевдоклас `:valid`, який змінює колір кордону на зелений, якщо введено дійсне значення.
1. Елемент `button` повинен мати псевдоклас `:hover`, який змінює колір фону при наведенні курсору.
1. Використайте псевдоклас `:checked` на `.radio-group input[type="radio"]`, щоб додати колір кордону, колір фону та тінь блоку, якщо радіокнопку вибрано.
1. Використайте псевдоклас `:checked` на радіокнопках, щоб змінити колір тексту відповідної мітки, якщо її вибрано.
1. Додайте псевдоклас `:first-of-type` до елемента `input`, щоб змінити стиль першого поля для введення. Наприклад, змініть кути на заокруглені.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Ви повинні мати елемент `div` з класом `container`.

```js
assert.exists(document.querySelector("div.container"));
```

Елемент `div` має містити елемент `form`.

```js
assert.exists(document.querySelector("div.container > form"));
```

Форма має містити елемент `input` з типом `text` та `id` зі значенням `name`, де користувач буде вводити повне ім’я.

```js
assert.exists(document.querySelector("div.container > form input#name[type='text']"));
```

Ви повинні мати ще один елемент `input` з типом `email` та `id` зі значенням `email`, де користувач буде вводити електронну адресу.

```js
assert.exists(document.querySelector("div.container > form input#email[type='email']"));
```

Форма має містити елемент `select` з `id` зі значенням `position` і декількома елементами `option`.

```js
assert.exists(document.querySelector("div.container > form select#position"));
assert.isNotEmpty(document.querySelectorAll("div.container > form select#position > option"))
```

Ви повинні мати елемент `fieldset` або `section` з класом `.radio-group`.

```js
assert.lengthOf(document.querySelectorAll('form fieldset.radio-group, form section.radio-group'), 1)
```

В межах `.radio-group` має бути група елементів `input` з типом `radio` та відповідними мітками для вибору зайнятості. Назвою цієї групи має бути `availability`.

```js
assert.isAtLeast(document.querySelectorAll("div.container > form .radio-group input[type='radio'][name='availability']").length, 2);
assert.isEmpty(document.querySelectorAll("div.container > form .radio-group input[type='radio']:not([name='availability'])"))
```

Ви повинні мати елемент `textarea` з `id` зі значенням `message`, де користувач буде вводити повідомлення.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form textarea#message"));
```

You should associate every `input` element with a `label` element.

```js
const els = document.querySelectorAll('input');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label); 
  assert.equal(label.tagName, "LABEL");
})
```

Ви повинні мати елемент `button` з типом `submit`, щоб користувач надіслав форму.

```js
assert.isNotEmpty(document.querySelectorAll("div.container > form button[type='submit']"));
```

Додайте псевдоклас `:focus` до елементів `input` та `textarea`, щоб колір кордону змінювався під час фокусування. Використайте селектор списку в зазначеному порядку.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input:focus, textarea:focus').getPropertyValue('border-color'))
```

Елементи `input`, `select` та `textarea` повинні мати псевдоклас `:invalid`, який змінює колір кордону на червоний, якщо введено недійсне значення. Використайте селектор списку в зазначеному порядку.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:invalid, select:invalid, textarea:invalid').getPropertyValue('border-color'), 'red')
```

Елементи `input`, `select` та `textarea` повинні мати псевдоклас `:valid`, який змінює колір кордону на зелений, якщо введено дійсне значення. Використайте селектор списку в зазначеному порядку.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input:valid, select:valid, textarea:valid').getPropertyValue('border-color'), 'green')
```

Елемент `button` повинен мати псевдоклас `:hover`, який змінює колір фону при наведенні курсору.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('button:hover').getPropertyValue('background-color'))
```

Використайте псевдоклас `:checked` на `.radio-group input[type="radio"]`, щоб додати колір кордону, якщо радіокнопку вибрано.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('border-color'));
```

Використайте псевдоклас `:checked` на `.radio-group input[type="radio"]`, щоб додати колір фону, якщо радіокнопку вибрано.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('background-color'));
```

Використайте псевдоклас `:checked` на `.radio-group input[type="radio"]`, щоб додати тінь блоку, якщо радіокнопку вибрано.

```js
const style = new __helpers.CSSHelp(document).getStyle('.radio-group input[type="radio"]:checked');
assert.isNotEmpty(style.getPropertyValue('box-shadow'));
```

Використайте псевдоклас `:checked` на радіокнопках, щоб змінити колір тексту відповідної мітки, якщо її вибрано.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isNotEmpty(els);
els.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.equal(label.tagName, "LABEL");
  input.checked = false;
  const colorBefore = getComputedStyle(label).color;
  input.checked = true;
  const colorAfter = getComputedStyle(label).color;
  assert.notEqual(colorAfter, colorBefore);
})
```

Додайте псевдоклас `:first-of-type` до елементів `input`, щоб змінити стиль першого поля для введення.

```js
const pseudo = new __helpers.CSSHelp(document).getStyle('input:first-of-type');
assert.exists(pseudo);
assert.isNotEmpty([...pseudo]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application Form</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Job Application</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="container">
        <h1>Job Application Form</h1>
        <form id="job-application">
            <label for="name">Full Name:</label>
            <input type="text" id="name" name="name" placeholder="Enter your name" required>

            <label for="email">Email:</label>
            <input type="email" id="email" name="email" placeholder="Enter your email" required>

            <label for="position">Position:</label>
            <select id="position" name="position" required>
                <option value="">Select a position</option>
                <option value="developer">Developer</option>
                <option value="designer">Designer</option>
                <option value="manager">Manager</option>
            </select>

            <fieldset class="radio-group">
              <legend>Availability:</legend>
                <input type="radio" id="fullTime" name="availability" value="fullTime" checked>
                <label for="fullTime">Full-Time</label>
                <input type="radio" id="partTime" name="availability" value="partTime">
                <label for="partTime">Part-Time</label>
            </fieldset>

            <label for="message">Why do you want this job?</label>
            <textarea id="message" name="message" rows="5" placeholder="Write your motivation" required></textarea>

            <button type="submit">Submit Application</button>
        </form>
    </div>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  background-color: white;
  padding: 20px;
}

.container {
  max-width: 600px;
  margin: 0 auto;
  background-color: white;
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
}

h1 {
  text-align: center;
}

label {
  display: block;
  margin-bottom: 5px;
  font-weight: bold;
}

input, select, textarea {
  width: 100%;
  padding: 10px;
  margin-bottom: 15px;
  border: 1px solid #ccc;
  border-radius: 4px;
  box-sizing: border-box;
}

.radio-group {
  display: flex;
  align-items: center;
  margin-top: 15px;
}
.radio-group input[type="radio"] {
  margin-right: 10px;
  -webkit-appearance: none;
  -moz-appearance: none;
  appearance: none;
  width: 20px;
  height: 20px;
  border: 2px solid #ccc;
  border-radius: 50%;
  outline: none;
  transition: border-color 0.2s;
}

.radio-group label {
  margin: 0;
  font-weight: normal;
  cursor: pointer;
}

button {
  width: 100%;
  padding: 10px;
  background-color: green;
  border: none;
  color: white;
  border-radius: 4px;
  font-size: 16px;
  cursor: pointer;
}

input:focus, textarea:focus {
  border-color: blue;
  outline: none;
}

input:invalid, select:invalid, textarea:invalid {
  border-color: red;
}

input:valid, select:valid, textarea:valid {
  border-color: green;
}

.radio-group input[type="radio"]:checked {
  border-color: green;
  background-color: green;
  box-shadow: inset 0 0 0 4px white;
}

button:hover {
  background-color: black;
}

input[type="radio"]:checked + label {
  color: green;
}

button:disabled {
  background-color: grey;
}

input:first-of-type {
  border-top-left-radius: 4px;
  border-top-right-radius: 4px;
}

textarea:nth-child(5) {
  border-bottom-left-radius: 4px;
  border-bottom-right-radius: 4px;
}
```
