---
id: 66f53dc2c5bd6a11d6c3282f
title: パスワードジェネレーターアプリを作成する
challengeType: 26
dashedName: lab-password-generator
---

# --description--

このラボでは、関数を使ってランダムなパスワードジェネレーターを作成する練習をします。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. 長さを示すパラメータを受け取る`generatePassword`という名前の関数を作成してください。パラメータ名は自由に決めてかまいません。
2. 関数はランダムに生成されたパスワードを表す文字列を返す必要があります。以下の文字列とさまざまな`Math`メソッドを使って、ランダムな文字を含む新しい文字列を返してください: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`。
3. `password`という変数を定義し、望むパスワードの長さを表す数値の引数で`generatePassword`関数を呼び出した結果を代入してください。
4. `console.log`で、メッセージ`Generated password:`と`password`変数を空白で区切って連結した単一の文字列をログに出力してください。

# --hints--

パラメータを持つ`generatePassword`関数を作成してください。パラメータ名は自由に決めてかまいません。

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

`generatePassword`関数は文字列を返す必要があります。

```js
const result = generatePassword(5);
assert.isString(result);
```

`generatePassword`関数は正しい長さの新しい文字列を返す必要があります。

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

関数は有効な文字を使ったランダムに生成されたパスワードを返す必要があります。

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

関数は複数の異なる文字を含むランダムに生成されたパスワードを返す必要があります。

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

関数は呼び出されるたびに新しいランダムな文字列を返す必要があります。

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

`password`変数を作成してください。

```js
assert.isDefined(password);
```

`password`変数は文字列である必要があります。

```js
assert.isString(password);
```

`generatePassword`関数を数値の引数で呼び出し、返されたパスワードを`password`変数に代入してください。

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

`Generated password:`と`password`を空白で区切って連結した単一の文字列を、`+`またはテンプレートリテラルを使ってログに出力してください。

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
