---
id: 6989f555728e63ad8b349d84
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Якщо поточний рядок ще не завершений, потрібно дослідити наступні можливі стани. Додайте умову `else` для поточної інструкції `if`.

Всередині блоку `else` ви будете обробляти додавання відкриваючих круглих дужок. Додайте інструкцію `if`, щоб перевірити, чи `opensUsed` менше за `pairs`. Це гарантує, що ви не використаєте більше відкриваючих дужок, ніж дозволено.

Якщо ця умова `true`, додайте новий масив у `queue`: `[current + '(', opensUsed + 1, closesUsed]`. Це представляє стан після додавання відкриваючої дужки.

# --hints--

У вас має бути умова `else` для вашої інструкції `if` всередині циклу `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

У вас має бути інструкція `if` всередині блоку `else` з умовою `opensUsed < pairs`.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Ви маєте додати `[current + '(', opensUsed + 1, closesUsed]` до черги, коли умова `true` виконується.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
