---
id: 6989eb65324d97775c4965fe
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте функцію, яка генерує всі допустимі комбінації круглих дужок, використовуючи підхід пошуку у ширину (BFS). Наприклад, допустимі комбінації двох пар круглих дужок — це `(())` та `()()`.

Почніть зі створення функції з назвою `genParentheses` з одним параметром `pairs`. Поки що поверніть з функції порожній масив.

# --hints--

Ви повинні визначити функцію з назвою `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Ваша функція `genParentheses` має мати один параметр з назвою `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Ваша функція має повертати порожній масив.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
