---
id: 6989f207306332c06b0add39
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Для підходу пошуку у ширину ви будете використовувати чергу, щоб відстежувати різні стани під час побудови комбінацій круглих дужок. Кожен стан буде представлений як масив, що містить три елементи:

- Поточний рядок, який будується
- Кількість відкриваючих круглих дужок, використаних до цього моменту
- Кількість закриваючих круглих дужок, використаних до цього моменту

Створіть змінну з назвою `queue` і ініціалізуйте її масивом, що містить один масив: `['', 0, 0]`. Це представляє початковий стан з порожнім рядком і нульовою кількістю використаних дужок.

# --hints--

Ви повинні оголосити змінну з назвою `queue` у вашій функції `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

Ви повинні ініціалізувати `queue` масивом, що містить масив `['', 0, 0]`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
