---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Тепер потрібно обробити додавання закриваючих круглих дужок. Головне правило — ви можете додати закриваючу дужку лише якщо це зберігає баланс, тобто має бути більше відкриваючих дужок, ніж закриваючих.

Все ще всередині блоку `else` додайте другу інструкцію `if`, щоб перевірити, чи `closesUsed` менше за `opensUsed`.

Якщо ця умова істинна, додайте ще один новий масив у `queue`: `[current + ')', opensUsed, closesUsed + 1]`. Це представляє стан після додавання закриваючої дужки.

# --hints--

У вас має бути друга інструкція `if` з умовою `closesUsed < opensUsed` всередині вашого блоку `else`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Ви маєте додати `[current + ')', opensUsed, closesUsed + 1]` у чергу, коли умова істинна.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
