---
id: 6989ef885db3a3488a13b426
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Перед тим як реалізувати основний алгоритм, потрібно перевірити вхідні дані. Параметр `pairs` має бути цілим числом, оскільки він представляє кількість пар круглих дужок, які потрібно згенерувати.

Додайте інструкцію `if` на початку вашої функції, щоб перевірити, чи `pairs` не є цілим числом. Для цього використайте метод `Number.isInteger()`.

Якщо умова істинна, поверніть рядок `The number of pairs should be an integer`.

# --hints--

У вас має бути оператор if з умовою `!Number.isInteger(pairs)` всередині функції `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Ви маєте повертати рядок, коли вхідні дані не є цілим числом.

```js
let result = genParentheses("i");
assert.isString(result);
```

Ви маєте повертати рядок `The number of pairs should be an integer`, коли вхідні дані є рядком.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
