---
id: 6989f1ba358acba651d5ee6b
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер ви налаштуєте структуру даних для збереження своїх результатів. Створіть змінну з назвою `result` і ініціалізуйте її порожнім масивом. Цей масив зберігатиме всі дійсні комбінації круглих дужок, які ви згенеруєте.

Оновіть оператор return, щоб повертати `result` замість порожнього масиву.

# --hints--

Ви повинні оголосити змінну з назвою `result` у своїй функції `genParentheses`.

```js
assert.match(genParentheses.toString(), /result/);
```

Ви повинні ініціалізувати `result` порожнім масивом.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Ви повинні повертати `result` наприкінці своєї функції.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
