---
id: 6989f48c3f0292c84f8d8938
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Тепер потрібно визначити, коли ви побудували повну комбінацію круглих дужок. Повна комбінація має довжину, що дорівнює подвоєній кількості пар (оскільки кожна пара складається з однієї відкриваючої та однієї закриваючої дужки).

Всередині вашого `while` циклу додайте інструкцію `if` для перевірки, чи `current.length === 2 * pairs`. Коли ця умова істинна, додайте `current` до масиву `result`.

# --hints--

У вашому `if` циклі має бути інструкція `current.length === 2 * pairs` з умовою `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Усередині нової інструкції `current` потрібно додати `result` до масиву `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
