---
id: 6989eb65324d97775c4965fe
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте функцію, яка генерує всі допустимі комбінації круглих дужок, використовуючи підхід пошуку у ширину (BFS). Наприклад, допустимі комбінації двох пар круглих дужок — це `(())` та `()()`.

Почніть зі створення функції з назвою `genParentheses` з одним параметром `pairs`. Поки що поверніть з функції порожній масив.

# --hints--

Ви повинні визначити функцію з назвою `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
```

Ваша функція `genParentheses` має мати один параметр з назвою `pairs`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.equal(genParentheses?.parameters.length, 1);
assert.equal(genParentheses?.parameters[0].toString(), 'pairs');
```

Ваша функція має повертати порожній масив.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.genParentheses);
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
