---
id: 6989f10a72ad424d4ad6d514
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Далі потрібно перевірити, що число пар принаймні одне, адже не можна згенерувати комбінації круглих дужок із нульовою або від’ємною кількістю пар.

Додайте ще одну `if` інструкцію, щоб перевірити, чи `pairs` менше за `1`. Якщо ця умова істинна, поверніть рядок `The number of pairs should be at least 1`.

# --hints--

У вас має бути друга `if` інструкція з умовою `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Ви повинні повертати рядок `The number of pairs should be at least 1`, коли `pairs` менше за `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
