---
id: 6989f2e35e6941e13e445156
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер ви реалізуєте основний цикл BFS. Створіть `while` цикл, який триватиме, доки `queue` не порожній, тобто доки `queue.length` більше за `0`.

Всередині циклу виводьте `queue` у консоль.

# --hints--

Ви маєте створити `while` цикл з умовою `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Ви маєте виводити `queue` у консоль всередині вашого `while` циклу.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
