---
id: 6989f3632f8b8bdf168aa321
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Усередині вашого циклу `while` використайте `queue.shift()`, щоб видалити та отримати перший елемент із черги. Це реалізує поведінку «перший прийшов — перший пішов» (FIFO), характерну для BFS.

Розпакуйте цей масив у три змінні: `current`, `opensUsed` та `closesUsed`.

# --hints--

Ви повинні використати `queue.shift()`, щоб отримати перший елемент із черги.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Результат слід розпакувати у три змінні: `current`, `opensUsed` та `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
