---
id: 6989f788362592a7238b1f94
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Тепер вам більше не потрібно виводити чергу в консоль. Тож видаліть `console.log(queue)` зі свого `while` циклу.

# --hints--

У вашому `console.log(queue)` циклі не має бути `while`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
--fcc-editable-region--
    console.log(queue);
--fcc-editable-region--

    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
      if (closesUsed < opensUsed) {
        queue.push([current + ')', opensUsed, closesUsed + 1]);
      }
    }
  }
  return result;
}

console.log(genParentheses(2));
```
