---
id: 686d749d6f32f708fdf2b7ad
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У попередньому уроці ви дізналися про елементи `svg` і як їх часто використовують як іконки в проєктах. У реальних кодових базах зазвичай покладаються на бібліотеки іконок, тому не потрібно створювати елементи `svg` з нуля. Проте в цьому практичному занятті ви створите іконку серця, щоб вивчити основні атрибути, які використовуються всередині елемента `svg`.

Почніть зі створення елемента `svg` на сторінці.

# --hints--

У вас має бути елемент `svg` на сторінці.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
