---
id: 60f02e7361b68405e27b62a5
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop lernen Sie, wie Sie Formulare gestalten, indem Sie ein Registrierungsformular entwerfen.

Der gesamte HTML-Boilerplate-Code (`DOCTYPE`, `html`, `head` und `body`) wurde für Sie bereitgestellt.

Fügen Sie innerhalb des `body` eine Überschrift hinzu, um dem Formular Kontext zu geben, indem Sie ein `h1`-Element mit dem Text `Registration Form` verwenden.

# --hints--

Sie sollten das `h1` innerhalb des `body` hinzufügen.

```js
assert.exists(document.querySelector('body > h1'));
```

Sie sollten dem `h1` den Text `Registration Form` geben.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
