---
id: 60f030d388cb74067cf291c3
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Unter der Überschrift verwenden Sie den folgenden Text in einem Absatz-Element, um die Nutzer zur Registrierung zu ermutigen:

```md
Please fill out this form with the required information
```

# --hints--

Sie sollten ein `p`-Element innerhalb des `body` hinzufügen.

```js
assert.exists(document.querySelector('body > p'));
```

Sie sollten das `p`-Element unterhalb des `h1` hinzufügen.

```js
assert.exists(document.querySelector('h1 + p'));
```

Sie sollten dem `p`-Element den Text `Please fill out this form with the required information` geben.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```
