---
id: 60f1a9cbd23023082e149fee
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Das ist besser. Machen Sie nun den Hintergrund augenfreundlich, indem Sie die `background-color` von `body` zu `#1b1b32` ändern. Um den Text sehen zu können, ändern Sie dann die `color` zu `#f5f6f7`.

# --hints--

Sie sollten die `background-color` innerhalb des `body`-Elementselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Sie sollten der `background-color` einen Wert von `#1b1b32` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Sie sollten die `color` innerhalb des `body`-Elementselektors hinzufügen.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Sie sollten der `color` einen Wert von `#f5f6f7` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
