---
id: 60f1922fcbd2410527b3bd89
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit Viewport-Einheiten arbeitet.

Denken Sie daran, dass die `vh`-Einheit für die Viewport-Höhe steht und 1 % der `height` des Viewports entspricht. Dadurch ist sie relativ zur Viewport-Höhe.

Es ist Zeit, das Projekt mit etwas CSS aufzuwerten. Beginnen Sie damit, dem `body` eine `width` von `100%` und eine `height` von `100vh` zu geben.

# --hints--

Sie sollten den `body`-Elementselektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Sie sollten dem `body` eine `width` von `100%` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Sie sollten dem `body` eine `height` von `100vh` geben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
