---
id: 60f5d2776c854e069560fbe6
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Das erste `fieldset` enthält die Felder für Name, E-Mail und Passwort. Beginnen Sie damit, vier `label`-Elemente zum ersten `fieldset` hinzuzufügen.

# --hints--

Es sollten insgesamt drei `fieldset`-Elemente vorhanden sein.

```js
assert.equal(document.querySelectorAll('fieldset')?.length, 3);
```

Die `fieldset`-Elemente sollten alle direkte Kinder des `form`-Elements sein.

```js
assert.equal(document.querySelectorAll('form > fieldset')?.length, 3);
```

Die vier `label`-Elemente sollten alle direkte Kinder des ersten `fieldset`-Elements sein. Achten Sie darauf, die `label`-Elemente zu schließen.

```js
assert.equal(document.querySelectorAll('form fieldset:first-child > label')?.length, 4);
```

Die letzten beiden `fieldset`-Elemente sollten leer sein.

```js
const LastTwoFieldsetElements = Array.from(document.querySelectorAll("form fieldset:not(fieldset:first-child)"));
assert.isTrue(LastTwoFieldsetElements.filter((ele) => ele.innerHTML?.replace(/\s/g, ""))?.length === 0);
```

Achten Sie darauf, die `label`-Elemente zu schließen.

```js
assert.lengthOf(document.querySelector('fieldset')?.innerHTML?.match(/<\/label>/g), 4)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
