---
id: 60f5c3e399ff1a05629964e4
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Wie der Titel schon andeutet, erstellen Sie ein Formular. Fügen Sie also nach dem `p`-Element ein `form` mit einem `action`-Attribut ein, das auf `https://register-demo.freecodecamp.org` zeigt.

# --hints--

Sie sollten ein `form`-Element neben dem `p`-Element hinzufügen.

```js
assert.exists(document.querySelector('p + form'));
```

Sie sollten dem `form` ein `action`-Attribut geben.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Sie sollten dem `action` einen Wert von `https://register-demo.freecodecamp.org` geben.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

Ihr `form`-Element sollte ein schließendes Tag `</form>` haben.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
