---
id: 60f805f813eaf2049bc2ceea
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Verschachteln Sie ein `input`-Element innerhalb jedes `label`. Achten Sie darauf, jedes `input` nach dem `label`-Text hinzuzufügen und nach dem Doppelpunkt ein Leerzeichen einzufügen.

# --hints--

Sie sollten vier `input`-Elemente zum `fieldset`-Element hinzufügen.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Sie sollten die `input`-Elemente innerhalb der `label`-Elemente verschachteln.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Sie sollten das erste `input` nach dem `label`-Text `Enter Your First Name:` hinzufügen und nach dem Doppelpunkt ein Leerzeichen einfügen.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

Sie sollten das zweite `input` nach dem `label`-Text `Enter Your Last Name:` hinzufügen und nach dem Doppelpunkt ein Leerzeichen einfügen.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

Sie sollten das dritte `input` nach dem `label`-Text `Enter Your Email:` hinzufügen und nach dem Doppelpunkt ein Leerzeichen einfügen.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

Sie sollten das vierte `input` nach dem `label`-Text `Create a New Password:` hinzufügen und nach dem Doppelpunkt ein Leerzeichen einfügen.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
