---
id: 60f803d5241e6a0433a523a1
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit `rem` Einheiten arbeitet. Denken Sie daran, dass die `rem` Einheit für die Wurzel-`em` steht und relativ zur Schriftartgröße des `html` Elements ist.

Da `label` Elemente standardmäßig inline sind, werden sie alle nebeneinander in derselben Zeile angezeigt, was das Lesen ihres Texts erschwert.

Um sie auf separaten Zeilen erscheinen zu lassen, fügen Sie dem `display: block` Element `label` hinzu und fügen Sie einen `margin` von `0.5rem 0` hinzu, um sie voneinander zu trennen.

# --hints--

Sie sollten den `label` Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Sie sollten eine `display` Eigenschaft mit dem Wert `block` hinzufügen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Sie sollten eine `margin` Eigenschaft mit dem Wert `0.5rem 0` hinzufügen.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
