---
id: 60fab8367d35de04e5cb7929
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Fügen Sie ein Ankerelement mit dem Text `Read our terms and conditions` vor dem neu hinzugefügten `label` ein. Setzen Sie das `href` auf:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

Fügen Sie dann den Text `I accept the terms and conditions` unmittelbar nach dem Eingabeelement in der Bezeichnung hinzu.

# --hints--

Sie sollten ein `a`-Element direkt vor dem `label` mit dem Text `Read our terms and conditions` haben.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.innerText.trim().replaceAll(/\s+/g, ' '), /Read our terms and conditions\.?/i);
```

Das `a`-Element sollte eine `href` von `https://www.freecodecamp.org/news/terms-of-service/` haben.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

Sie sollten ein `input`-Element innerhalb des `label` haben.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + a + label > input'));
```

Innerhalb des `label`, unmittelbar nach dem `input`-Element, sollten Sie den Text `I accept the terms and conditions` haben.

```js
assert.include(document.querySelector('fieldset:nth-child(3) + a + label')?.textContent.trim().replaceAll(/\s+/g, ' '), 'I accept the terms and conditions');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
--fcc-editable-region--
    
--fcc-editable-region--
    <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
