---
id: 60fabf0dd4959805dbae09e6
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

Fügen Sie nach dem ersten ein weiteres `label` mit dem Text `Input your age (years): ` hinzu. Verschachteln Sie dann ein `input` mit dem `type` von `number`.

Fügen Sie als Nächstes dem `min` ein `input`-Attribut mit dem Wert `13` hinzu, da sich Benutzer unter 13 Jahren nicht registrieren sollten. Außerdem werden Benutzer wahrscheinlich nicht älter als 120 Jahre sein; fügen Sie ein `max`-Attribut mit dem Wert `120` hinzu.

Wenn nun jemand versucht, das Formular mit Werten außerhalb dieses Bereichs abzusenden, erscheint eine Warnung, und das Formular wird nicht abgesendet. Probieren Sie es aus.

# --hints--

Sie sollten dem dritten `label` ein `fieldset` hinzufügen, nach dem vorhandenen `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

Sie sollten dem `label` den Text `Input your age (years): ` geben.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

Sie sollten dem `label` ein `input` mit `type` von `number` geben.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

Sie sollten dem `input` ein `min`-Attribut mit dem Wert `13` geben.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

Sie sollten dem `input` ein `max`-Attribut mit dem Wert `120` geben.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
