---
id: 60facde2d0dc61085b41063f
title: Step 34
challengeType: 0
dashedName: step-34
---

# --description--

Das `textarea`-Element verhält sich wie ein `input`-Element des Typs `text`, bietet jedoch den zusätzlichen Vorteil, mehrzeiligen Text sowie eine anfängliche Anzahl von Textzeilen und -spalten aufnehmen zu können.

Benutzende sollen sich mit einer Biografie registrieren können. Fügen Sie am Ende des `label` ein `Provide a bio:` mit dem Text `fieldset` hinzu. Fügen Sie ein `textarea`-Element innerhalb des `label`-Elements ein. Beachten Sie, dass das `textarea` ein schließendes Tag benötigt.

# --hints--

Sie sollten am Ende des dritten `label` ein `fieldset`-Element hinzufügen, nach den vorhandenen `label`-Elementen.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)'));
```

Sie sollten dem `label` den Text `Provide a bio:` geben.

```js
assert.match(document.querySelector('fieldset:nth-child(3) > label:nth-child(4)')?.innerText, /Provide a bio/);
```

Sie sollten ein `textarea`-Element innerhalb des `label` verschachteln.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(4) > textarea'));
```

Sie sollten dem `textarea`-Element ein öffnendes und ein schließendes Tag geben.

```js
assert.match(code, /<textarea\s*>[\s\S]*<\/textarea\s*>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
      <label>Input your age (years): <input type="number" min="13" max="120" /></label>
      <label>How did you hear about us?
        <select>
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
