---
id: 60fab4a123ce4b04526b082b
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

Derzeit können Nutzer das Formular absenden, ohne die Radio-Inputs auszuwählen. Obwohl Sie zuvor das `required`-Attribut verwendet haben, um anzugeben, dass eine Eingabe erforderlich ist, funktioniert das in diesem Fall nicht, weil das Hinzufügen von `required` zu beiden Eingaben den Nutzenden falsche Informationen vermittelt.

Um das zu lösen, können Sie Kontext darüber geben, was benötigt wird, indem Sie ein `legend`-Element mit dem Text `Account type (required)` vor den `label`-Elementen innerhalb des zweiten `fieldset` hinzufügen. Dann fügen Sie das `checked`-Attribut zum `Personal`-Input hinzu, um sicherzustellen, dass das Formular mit den erforderlichen Daten abgesendet wird.

# --hints--

Ihre `input`-Elemente sollten `type` von `radio` bleiben.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

Sie sollten ein `legend`-Element innerhalb des zweiten `fieldset`-Elements hinzufügen.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

Sie sollten `Account type (required)`-Text zum `legend`-Element hinzufügen.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Sie sollten dem `checked`-Input ein Attribut `Personal` geben.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
