---
id: 6197f667297bb30a552ce017
title: 步骤 9
challengeType: 0
dashedName: step-9
---

# --description--

由于 `.ground` 元素将在页面布局的堆叠中位于第三位，因此将其 `z-index` 设置为 `3`，并将其 `position` 设置为 `absolute`。

# --hints--

你应该将 `.ground` 的 `z-index` 设置为 `3`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.zIndex, '3');
```

你应该将 `.ground` 的 `position` 设置为 `absolute`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.position, 'absolute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));

}
--fcc-editable-region--
```
