---
id: 689f7996426c1534fa5bea66
title: Jenga mtafsiri wa alama za Golf
challengeType: 26
dashedName: build-a-golf-score-translator
---

# --description--

Katika mchezo wa Golf, kila shimo lina `par`, ikimaanisha wastani wa nambari za `strokes` ambazo mchezaji wa golf anatarajiwa kuzifanikisha ili kuingiza mpira kwenye shimo na kumaliza mchezo. Kulingana na umbali wa juu au chini ya `par` alama zako za `strokes` zinakuwa na jina la utani tofauti.

Katika maabara hii, utaandika kitendakazi kinachobadilisha `par` na `strokes` kuwa jina lao la utani.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `golfScore`.
1. `golfScore` inapaswa kuchukua hoja mbili za nambari, ambazo ni par ya kozi na idadi ya michoro iliyofanywa.
1. `golfScore` inapaswa kurudisha mfuatano wa herufi.
1. `golfScore` inapaswa kurudisha `"Hole-in-one!"` ikiwa `strokes` ni `1`.
1. `golfScore` inapaswa kurudisha `"Eagle"` ikiwa `strokes` ni ndogo au sawa na `par` pungufu `2`.
1. `golfScore` inapaswa kurudisha `"Birdie"` ikiwa `strokes` ni sawa na `par` pungufu `1`.
1. `golfScore` inapaswa kurudisha `"Par"` ikiwa `strokes` ni sawa na `par`.
1. `golfScore` inapaswa kurudisha `"Bogey"` ikiwa `strokes` ni sawa na `par` jumlisha `1`.
1. `golfScore` inapaswa kurudisha `"Double Bogey"` ikiwa `strokes` ni sawa na `par` jumlisha `2`.
1. `golfScore` inapaswa kurudisha `"Go Home!"` ikiwa `strokes` ni kubwa au sawa na `par` jumlisha `3`.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `golfScore`.

```js
assert.isFunction(golfScore);
```

`golfScore` inapaswa kuchukua vigezo viwili.

```js
assert.lengthOf(golfScore, 2);
```

`golfScore` inapaswa kurudisha mfuatano wa herufi

```js
assert.isString(golfScore(4,1)); 
```

`golfScore(1, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert.strictEqual(golfScore(1,1), 'Hole-in-one!');
```

`golfScore(3, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert.strictEqual(golfScore(3,1), 'Hole-in-one!');
```

`golfScore(4, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert.strictEqual(golfScore(4,1), 'Hole-in-one!');
```

`golfScore(5, 1)` inapaswa kurudisha mfuatano wa herufi `Hole-in-one!`

```js
assert.strictEqual(golfScore(5,1), 'Hole-in-one!');
```

`golfScore(4, 2)` inapaswa kurudisha mfuatano wa herufi `Eagle`

```js
assert.strictEqual(golfScore(4,2), 'Eagle');
```

`golfScore(5, 2)` inapaswa kurudisha mfuatano wa herufi `Eagle`

```js
assert.strictEqual(golfScore(5,2), 'Eagle');
```

`golfScore(3, 2)` inapaswa kurudisha mfuatano wa herufi `Birdie`

```js
assert.strictEqual(golfScore(3,2), 'Birdie');
```

`golfScore(4, 3)` inapaswa kurudisha mfuatano wa herufi `Birdie`

```js
assert.strictEqual(golfScore(4,3), 'Birdie');
```

`golfScore(5, 4)` inapaswa kurudisha mfuatano wa herufi `Birdie`

```js
assert.strictEqual(golfScore(5,4), 'Birdie');
```

`golfScore(3, 3)` inapaswa kurudisha mfuatano wa herufi `Par`

```js
assert.strictEqual(golfScore(3,3), 'Par');
```

`golfScore(4, 4)` inapaswa kurudisha mfuatano wa herufi `Par`

```js
assert.strictEqual(golfScore(4,4), 'Par');
```

`golfScore(5, 5)` inapaswa kurudisha mfuatano wa herufi `Par`

```js
assert.strictEqual(golfScore(5,5), 'Par');
```

`golfScore(3, 4)` inapaswa kurudisha mfuatano wa herufi `Bogey`

```js
assert.strictEqual(golfScore(3,4), 'Bogey');
```

`golfScore(4, 5)` inapaswa kurudisha mfuatano wa herufi `Bogey`

```js
assert.strictEqual(golfScore(4,5), 'Bogey');
```

`golfScore(5, 6)` inapaswa kurudisha mfuatano wa herufi `Bogey`

```js
assert.strictEqual(golfScore(5,6), 'Bogey');
```

`golfScore(3, 5)` inapaswa kurudisha mfuatano wa herufi `Double Bogey`

```js
assert.strictEqual(golfScore(3,5), 'Double Bogey');
```

`golfScore(4, 6)` inapaswa kurudisha mfuatano wa herufi `Double Bogey`

```js
assert.strictEqual(golfScore(4,6), 'Double Bogey');
```

`golfScore(5, 7)` inapaswa kurudisha mfuatano wa herufi `Double Bogey`

```js
assert.strictEqual(golfScore(5,7), 'Double Bogey');
```

`golfScore(3, 7)` inapaswa kurudisha mfuatano wa herufi `Go Home!`

```js
assert.strictEqual(golfScore(3,7), 'Go Home!');
```

`golfScore(4, 8)` inapaswa kurudisha mfuatano wa herufi `Go Home!`

```js
assert.strictEqual(golfScore(4,8), 'Go Home!');
```

`golfScore(5, 9)` inapaswa kurudisha mfuatano wa herufi `Go Home!`

```js
assert.strictEqual(golfScore(5,9), 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
