---
id: 66f17db06803d11a1bd19a20
title: Crie um jogo de questionário
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar um array chamado `questions`.
1. O array `questions` deve conter pelo menos cinco objetos, cada um com as chaves `category`, `question`, `choices` e `answer`.
1. A chave `category` deve ter o valor de uma string que representa uma categoria de pergunta.
1. A chave `question` deve ter o valor de uma string que representa uma pergunta.
1. A chave `choices` deve conter o valor de um array contendo três strings, que são respostas alternativas para a pergunta.
1. A chave `answer` deve ter o valor de uma string, representando a resposta correta para a pergunta. Além disso, o valor de `answer` deve estar incluído no array `choices`.
1. Você deve ter uma função chamada `getRandomQuestion` que recebe um array de perguntas como parâmetro e retorna um objeto de pergunta aleatória do array.
1. Você deve ter uma função chamada `getRandomComputerChoice` que recebe o array das escolhas disponíveis como parâmetro e retorna uma resposta aleatória para a pergunta selecionada.
1. Você deve ter uma função chamada `getResults` que recebe o objeto da pergunta como o primeiro parâmetro e a escolha do computador como o segundo parâmetro. A função deve retornar `A escolha do computador está correta!` se a resposta estiver correta. Caso contrário, retorna `A escolha do computador está errada. A resposta correta é: <correct-answer>`, onde `<correct-answer>` é o valor da resposta correta para a pergunta escolhida.

# --hints--

Você deve criar um array chamado `questions`.

```js
assert.isArray(questions);
```

O array `questions` deve conter pelo menos cinco objetos, cada um com as chaves `category`, `question`, `choices` e `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

A chave `category` deve ter o valor de uma string que representa uma categoria de pergunta.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

A chave `question` deve ter o valor de uma string que representa uma pergunta.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

A chave `choices` deve ter o valor de um array contendo três strings diferentes entre si.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

A chave `answer` deve ter o valor de uma string.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

O valor de `answer` deve ser incluído no array `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Você deve ter uma função chamada `getRandomQuestion` que recebe um array de perguntas como parâmetro e retorna um objeto de pergunta aleatória do array.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Você deve ter uma função chamada `getRandomComputerChoice` que recebe o array das escolhas disponíveis como parâmetro e retorna uma resposta aleatória para a pergunta selecionada.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Você deve ter uma função chamada `getResults`.

```js
assert.isFunction(getResults);
```

Sua função `getResults` deve receber o objeto da pergunta como o primeiro parâmetro e a escolha do computador como o segundo parâmetro.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Se a escolha do computador corresponder à resposta, `getResults` deve retornar `A escolha do computador está correta!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Se a escolha do computador não corresponder à resposta, `getResults` deve retornar `A escolha do computador está errada.` A resposta correta é: <correct-answer>`, onde `<correct-answer>` é o valor da resposta correta para a pergunta escolhida.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

Sua função `getResults` deve usar comparação de igualdade exata, não correspondência de substring.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
