---
id: 699f3738921fb79a651ded56
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa ni wakati wa kuita kitendakazi chako.

Tengeneza kigezo kinachoitwa `result` na weka thamani yake kuwa wito wa kitendakazi `square` ukiwa na hoja ya `5`.

Chini ya hapo, andika thamani ya `result` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `result`.

```js
assert.isDefined(result);
```

Kigezo chako cha `result` kinapaswa kuwa nambari.

```js
assert.isNumber(result);
```

Kigezo chako cha `result` kinapaswa kuwa mraba wa `5`.

```js
assert.strictEqual(result, 25);
```

Unapaswa kuandika kigezo cha `result` kwenye konsoli.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
