---
id: 699f4861dd5b5f164e1bda4a
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Faida nyingine ya kutumia TypeScript ni kuhakikisha kwamba aina za kurudisha zinatekelezwa na ni zile unazotarajia. Wakati mwingine TypeScript inaweza kubaini aina ya kurudisha ya kitendakazi, lakini pia unaweza kuibainisha waziwazi.

Ili kuonyesha hili, badilisha aina ya kurudisha ya kitendakazi cha `getAverage` kuwa mfuatano wa herufi `"something"`. 

# --hints--

Kitendakazi chako cha `getAverage` kinapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Kitendakazi chako cha `getAverage` kinapaswa kurudisha mfuatano wa herufi `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
