---
id: 699f4ec939f10da62fb56c3e
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako.

Tengeneza kigezo kinachoitwa `raisedResult` na weka thamani yake kuwa matokeo ya kuita `raiseTo` kwa hoja `2` na `3`.

Kisha, andika `raisedResult` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Kigezo chako cha `raisedResult` kinapaswa kuwa nambari.

```js
assert.isNumber(raisedResult);
```

Unapaswa kuandika `raisedResult` kwenye konsoli.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
