---
id: 699f52220de88c8091032076
title: Hatua 18
challengeType: 1
dashedName: step-18
---

# --description--

Ukifungua konsoli, matokeo ni `NaN` kwa sababu hoja ya pili kwa `exponent` haikuwahi kutolewa. Kuna masasisho kadhaa yanayohitajika kwa kitendakazi cha `raiseTo` ambacho kitaondoa tatizo hili.

Mabadiliko ya kwanza ni kurudisha `base ** 2`.

Katika hatua inayofuata, utaona jinsi ya kushughulikia kigezo cha hiari cha `exponent`.

# --hints--

Kitendakazi chako cha `raiseTo` kinapaswa kurudisha nambari.

```js
assert.isNumber(raiseTo(4));
```

Kitendakazi chako cha `raiseTo` kinapaswa kurudisha thamani ya `base` iliyoinuliwa kwa nguvu ya `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
