---
id: 699f49669ea51a31d2ba89d8
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Kama unavyokumbuka kutoka kwa mafundisho ya awali, unaweza kuweka aina ya kurudi kwa kitendakazi waziwazi kama hii:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

Katika mfano huu, aina ya kurudi ya kitendakazi cha `circleArea` ni `number`. Ukijaribu kurudisha aina tofauti, TypeScript itaonyesha kosa.

Sahihisha kitendakazi chako cha `getAverage` ili kuweka aina ya kurudi waziwazi kuwa `number`. 

# --hints--

Kitendakazi chako cha `getAverage` kinapaswa kuwa na aina ya kurudi ya `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
