---
id: 699f1d72e906e4616155290d
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utaendelea kujifunza kuhusu maelezo ya aina na jinsi zinavyofanya kazi na vitendakazi kwa kujenga kifaa cha hesabu salama kwa aina.

Anza kwa kuunda kitendakazi kinachoitwa `square` kinachokubali kigezo na kurudisha mraba wa nambari.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Kitendakazi chako cha `square` kinapaswa kuwa na kigezo kimoja.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Kitendakazi chako cha `square` kinapaswa kurudisha mraba wa nambari.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
