---
id: 699f434e294a9e11cf4f65f0
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Sasa ni wakati wa kuongeza msaidizi mwingine wa hesabu katika mradi wako wa zana.

Anza kwa kuunda kitendakazi kinachoitwa `getAverage` kinachopokea kigezo kimoja. Kitendakazi chako kinapaswa kurudisha wastani wa safu ya data ya nambari.

Katika hatua chache zijazo, utaongeza maelezo ya aina kwa kitendakazi hiki.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `getAverage`.

```js
assert.isFunction(getAverage);
```

Kitendakazi chako kinapaswa kupokea kigezo kimoja.

```js
assert.strictEqual(getAverage.length, 1);
```

Kitendakazi chako kinapaswa kurudisha nambari.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Kitendakazi chako kinapaswa kurudisha matokeo sahihi kwa safu yoyote ya nambari.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
