---
id: 699f3cb8d057f29954205d7d
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Kama unavyokumbuka kutoka mafundoni ya awali, unaweza kuongeza maelezo ya aina kwa vigezo vya kitendakazi kama hivi:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

Katika mfano huu, TypeScript itaonyesha kosa ikiwa utajaribu kuita kitendakazi cha `add` kwa hoja ambazo si nambari.

Sahihisha kitendakazi chako cha `square` ili kujumuisha maelezo ya aina kwa kigezo cha `num`. Aina inapaswa kuwa `number`.

# --hints--

Unapaswa kuongeza maelezo ya aina kwa kigezo cha `num` katika kitendakazi cha `square`. Aina inapaswa kuwa `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
--fcc-editable-region--
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
