---
id: 699f52fac083214b6bf51de6
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Katika hatua hii ya mwisho, utahitaji kuongeza ukaguzi wa masharti kwa kigezo cha `exponent` katika kitendakazi cha `raiseTo`.

Ikiwa hoja ya `exponent` itatolewa, kitendakazi kinapaswa kurudisha `base` kilichoinuliwa kwa nguvu ya `exponent`. Ikiwa hoja ya `exponent` haitatolewa, kinapaswa kurudisha `base ** 2`.

Kwa hatua hiyo ya mwisho, kifaa chako cha hesabu chenye usalama wa aina kimekamilika!

# --hints--

Kitendakazi chako cha `raiseTo` kinapaswa kurudisha `base ** exponent` wakati hoja ya `exponent` itatolewa.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Kitendakazi chako cha `raiseTo` kinapaswa kurudisha `base ** 2` wakati hoja ya `exponent` haitatolewa.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }
  
  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
