---
id: 699f4a3e7b4c6bfa538cc4f7
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Ukifungua konsoli, utaona ujumbe wa kosa ufuatao:

```md
Type 'string' is not assignable to type 'number'.
```

Ili kutatua kosa hili, unahitaji kubadilisha return kurudi `sum / numbers.length`.

# --hints--

Kitendakazi chako cha `getAverage` kinapaswa kurudisha nambari.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Kitendakazi chako cha `getAverage` kinapaswa kurudisha wastani wa nambari zilizopo katika safu ya data.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
--fcc-editable-region--
  return "something";
--fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

