---
id: 699f3f636aa840b7201aa37c
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Ukifungua konsoli, unapaswa kuona ujumbe wa kosa wa TypeScript ufuatao:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Sasa kwamba kigezo kimeelezwa waziwazi kuwa nambari, TypeScript ina uwezo wa kugundua kosa la kujaribu kupitisha hoja ya mfuatano wa herufi kwa kitendakazi cha `square`. 

Badilisha hoja inayopitishwa kwa kitendakazi cha `square` kuwa nambari `5` ili kurekebisha kosa na kuona matokeo sahihi kwenye konsoli.

# --hints--

Wito wako wa kitendakazi cha `square` unapaswa kuwa na hoja ya `5` badala ya `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
