---
id: 699f5063d2393c674b1f45dc
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Kama unavyokumbuka kutoka hatua iliyopita, ikiwa unaita kitendakazi kwa idadi isiyo sahihi ya hoja, TypeScript itatupa kosa.

Kuna njia ya kuweka vigezo kuwa hiari kwa kuongeza `?` baada ya jina la kigezo. Hii inakuwezesha kuita kitendakazi kwa hoja chache kuliko vigezo.

Hapa kuna mfano:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

Katika mfano huu, ikiwa kigezo cha `greeting` hakitolewi wakati wa kuita kitendakazi cha `greet`, kitakuwa `undefined`, na kitendakazi kitarejesha salamu ya chaguo-msingi.

Sahihisha kitendakazi chako cha `raiseTo` ili kigezo cha `exponent` kiwe hiari.

# --hints--

Kitendakazi chako cha `raiseTo` kinapaswa kuwa na kigezo cha `exponent` kilichowekwa kuwa hiari.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
--fcc-editable-region--
  return base ** exponent;
}

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
