---
id: 699f38937b053e4679878ff1
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Unapofanya kazi na JavaScript ya kawaida, hakuna kinachokuzuia kufanya wito wa kitendakazi cha `square` kwa kutumia mfuatano wa herufi au safu ya data, jambo ambalo litasababisha matokeo yasiyotegemewa.

Ili kuonyesha hili, badilisha wito wa kitendakazi chako kutumia mfuatano wa herufi `"something"` badala ya nambari `5`.

Katika funzo lijalo, utajifunza jinsi ya kutumia TypeScript kuongeza usalama wa aina katika msimbo wako na kuzuia makosa kama haya.

# --hints--

Wito wako wa kitendakazi cha `square` unapaswa kuwa na hoja ya `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
