---
id: 699f513e01eb1768cc5df726
title: Hatua 17
challengeType: 1
dashedName: step-17
---

# --description--

Sasa kigezo cha `exponent` kimewekwa kama hiari, unaweza kuita kitendakazi kwa hoja moja tu.

Sahihisha wito wako wa kitendakazi cha `raiseTo` ili upitishe nambari `3` tu.

# --hints--

Kitendakazi chako cha `raiseTo` kinapaswa kuitwa kwa hoja moja tu ya `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
