---
id: 699f4639ae0016caaf81b7c6
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako.

Anza kwa kuunda kigezo kinachoitwa `avgResult` na ukaweke thamani yake kwa wito wa kitendakazi `getAverage`. Hoja ya wito huo wa kitendakazi inapaswa kuwa `[2, 14, 26, 8]`.

Kisha andika thamani ya `avgResult` kwenye konsoli.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `avgResult`.

```js
assert.isDefined(avgResult);
```

Unapaswa kuweka `getAverage([2, 14, 26, 8])` kwenye kigezo chako cha `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Kigezo chako cha `avgResult` kinapaswa kuwa nambari.

```js
assert.isNumber(avgResult);
```

Unapaswa kuandika `avgResult` kwenye konsoli.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
