---
id: 699f412114c87221f16a99e3
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Faida nyingine ya kutumia TypeScript ni kwamba itagundua makosa ikiwa utapita hoja nyingi sana au chache sana kwa wito wa kitendakazi.

Ili kuonyesha hili, sasisha wito wako wa kitendakazi cha `square` ili uchukue hoja mbili za `5` na `10`.

# --hints--

Wito wako wa kitendakazi cha `square` unapaswa kuwa na hoja mbili za `5` na `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
