---
id: 699f41b8da20f816b5c7dece
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ukifungua konsoli, unapaswa kuona ujumbe wa kosa ufuatao:

```md
Expected 1 arguments, but got 2.
```

Ili kutatua kosa hilo, sasisha wito wako wa kitendakazi `square` ili uchukue hoja moja tu ya `5`.

# --hints--

Wito wako wa kitendakazi `square` unapaswa kuwa na hoja moja tu ya `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
