---
id: 699f458e47fb88c756fdba81
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Kwa sasa, kigezo cha `numbers` kina aina ya `any` isiyoelezwa wazi. Ili kurekebisha hili, unaweza kuweka aina ya kigezo hicho wazi kama safu ya data ya nambari.

Kumbuka kutoka mafundisho ya awali, kwamba unaweza kuweka aina ya safu ya data ya nambari kama ifuatavyo:

```ts
const numbers: number[] = [1, 2, 3];
```

Sahihisha kitendakazi chako cha `getAverage` kwa kuweka aina wazi ya kigezo cha `numbers` kama safu ya data ya nambari.

# --hints--

Kitendakazi chako cha `getAverage` kinapaswa kuwa na kigezo kinachoitwa `numbers` ambacho kimeelezwa wazi kuwa safu ya data ya nambari.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
--fcc-editable-region--
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
