---
id: 663d18c574456976d11b6e87
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在遊戲中，將會有多個回合。第一個達到三分的人贏得遊戲。

在此步驟中，你將專注於判斷玩家是否贏得該回合。

完成 `hasPlayerWonTheRound` 函式（程式）。此函式有兩個參數：`player` 和 `computer`。如果玩家贏得該回合，函式應傳回 `true`，如果玩家輸掉或平手，則傳回 `false`。

以下是玩家贏得一回合的條件：

- 如果玩家選擇 `"Rock"` 且電腦選擇 `"Scissors"`
- 如果玩家選擇 `"Scissors"` 且電腦選擇 `"Paper"`
- 如果玩家選擇 `"Paper"` 且電腦選擇 `"Rock"`

已為你提供了幾個函式呼叫來測試你的函式。

# --hints--

你的 `hasPlayerWonTheRound` 函式（程式）應該傳回一個布林值。

```js
assert.isBoolean(hasPlayerWonTheRound("Rock", "Scissors"));
```

如果玩家選擇了 `"Rock"` 且電腦選擇了 `"Scissors"`，你的 `hasPlayerWonTheRound` 函式（程式）應該傳回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Rock", "Scissors"));
```

如果玩家選擇了 `"Scissors"` 且電腦選擇了 `"Paper"`，你的 `hasPlayerWonTheRound` 函式（程式）應該傳回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Scissors", "Paper"));
```

你的 `hasPlayerWonTheRound` 函式（程式）應該在玩家選擇 `"Paper"` 且電腦選擇 `"Rock"` 時傳回 `true`。

```js
assert.isTrue(hasPlayerWonTheRound("Paper", "Rock"));
```

如果玩家和電腦選擇了相同的選項，你的 `hasPlayerWonTheRound` 函式（程式）應該傳回 `false`。

```js
assert.isFalse(hasPlayerWonTheRound("Rock", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Scissors", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Paper"));
```

你的 `hasPlayerWonTheRound` 函式（程式）如果電腦贏得該回合，應該傳回 `false`。

```js
assert.isFalse(hasPlayerWonTheRound("Scissors", "Rock"));
assert.isFalse(hasPlayerWonTheRound("Paper", "Scissors"));
assert.isFalse(hasPlayerWonTheRound("Rock", "Paper"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Rock, Paper, Scissors game</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Let's play Rock, Paper, Scissors!</h1>
    <main>
      <details class="rules-container">
        <summary>Rules to the game</summary>

        <p>You will be playing against the computer.</p>
        <p>You can choose between Rock, Paper, and Scissors.</p>
        <p>The first one to three points wins.</p>

        <p>Here are the rules to getting a point in the game:</p>
        <ul>
          <li>Rock beats Scissors</li>
          <li>Scissors beats Paper</li>
          <li>Paper beats Rock</li>
        </ul>
        <p>
          If the player and computer choose the same option (Ex. Paper and
          Paper), then no one gets the point.
        </p>
      </details>

      <div class="score-container">
        <strong
          >Player Score: <span class="score" id="player-score">0</span></strong
        >
        <strong
          >Computer Score:
          <span class="score" id="computer-score">0</span></strong
        >
      </div>

      <section class="options-container">
        <h2>Choose an option:</h2>
        <div class="btn-container">
          <button id="rock-btn" class="btn">Rock</button>
          <button id="paper-btn" class="btn">Paper</button>
          <button id="scissors-btn" class="btn">Scissors</button>
        </div>
      </section>

      <div class="results-container">
        <p id="results-msg"></p>
        <p id="winner-msg"></p>
        <button class="btn" id="reset-game-btn">Play again?</button>
      </div>
    </main>
    <script src="./script.js"></script>
  </body>
</html>

```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --very-dark-blue: #0a0a23;
  --white: #ffffff;
  --yellow: #f1be32;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--very-dark-blue);
  text-align: center;
  color: var(--white);
}

h1 {
  margin: 15px 0 20px;
}

.btn {
  cursor: pointer;
  width: 100px;
  margin: 10px;
  color: var(--very-dark-blue);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

.rules-container {
  padding: 10px 0;
  margin: auto;
  border-radius: 15px;
  border: 5px solid var(--yellow);
  background-color: var(--white);
  color: var(--very-dark-blue);
}

.rules-container ul {
  list-style-type: none;
}

.rules-container p {
  margin: 10px 0;
}

@media (min-width: 760px) {
  .rules-container {
    width: 60%;
  }
}

.score-container {
  display: flex;
  justify-content: space-around;
  margin: 30px 0;
  font-size: 1.2rem;
}

.score {
  font-weight: 500;
}

.results-container {
  font-size: 1.3rem;
  margin: 15px 0;
}

#winner-msg {
  margin-top: 25px;
}

#reset-game-btn {
  display: none;
  margin: 20px auto;
}

```

```js
function getRandomComputerResult() {
  const options = ["Rock", "Paper", "Scissors"];
  const randomIndex = Math.floor(Math.random() * options.length);
  return options[randomIndex];
}

--fcc-editable-region--
function hasPlayerWonTheRound(player, computer) {
  
}

console.log(hasPlayerWonTheRound("Rock", "Scissors")); 
console.log(hasPlayerWonTheRound("Scissors", "Rock")); 
--fcc-editable-region--
```
