---
id: 5f332b23c2045fb843337579
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

To let visitors know the cafe was founded in `2020`, add a `p` element below the `h1` element with the text `Est. 2020`.

# --hints--

Du solltest ein öffnendes `<p>`-Tag haben.

```js
assert(code.match(/<p>/i));
```

Du solltest ein schließendes `</p>` Tag haben.

```js
assert(code.match(/<\/p>/i));
```

Du solltest dein vorhandenes `h1`-Element nicht ändern. Stelle sicher, dass du das abschließende Tag nicht gelöscht hast.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Dein `p`-Element sollte unter deinem `h1`-Element stehen.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Your `p` element should have the text `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

