---
id: 5f3313e74582ad9d063e3a38
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. The `title` element's text should be `Cafe Menu`.

# --hints--

Du solltest ein öffnendes `<head>`-Tag haben.

```js
assert.match(code, /<head>/i);
```

Du solltest ein schließendes `</head>`-Tag haben.

```js
assert.match(code, /<head>/i);
```

Du solltest ein öffnendes `<title>` Tag haben.

```js
assert.match(code, /<title>/i);
```

Du solltest ein schließendes `</title>`-Tag haben.

```js
assert.match(code, /<\/title>/i);
```

Dein `title`-Element sollte sich innerhalb deines `head`-Elements befinden.

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

Dein `title`-Element sollte den Text `Cafe Menu` haben. Möglicherweise musst du deine Rechtschreibung überprüfen.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
