---
id: 5f332a88dc25a0fd25c7687a
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

Du solltest ein öffnendes `<h1>`-Tag haben.

```js
assert.match(code, /<h1>/i);
```

Du solltest ein schließendes `</h1>`-Tag haben.

```js
assert.match(code, /<\/h1>/i);
```

Du solltest dein bestehendes `main`-Element nicht ändern.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Dein `h1`-Element sollte in deinem `main`-Element verschachtelt sein.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Dein `h1`-Element sollte den Text `CAMPER CAFE` enthalten.

```js
assert.match(code, /<h1>CAMPER CAFE<\/h1>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

