---
id: 5f344f9c805cd193c33d829c
title: Schitt 11
challengeType: 0
dashedName: step-11
---

# --description--

In previous lectures, you learned how to add `CSS` properties and values like this:

```css
element {
 property: value;
}
```

Zentriere den Inhalt des `h1`-Elements, indem du die `text-align`-Eigenschaft auf den Wert `center` setzt.

# --hints--

Du solltest ein `h1`-Selektor in deinem `style`-Element haben.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

Deine `text-align`-Eigenschaft sollte den Wert `center` haben.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

Dein `h1`-Selektor sollte die `text-align`-Eigenschaft auf `center` setzen.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

