---
id: 5f3477ae9675db8bb7655b30
title: Schritt 12
challengeType: 0
dashedName: step-12
---

# --description--

In the previous step, you used a <dfn>type selector</dfn> to style the `h1` element. Center the content of the `h2` and the `p` elements by adding a new type selector for each one to the existing `style` element.

# --hints--

Du solltest den vorhandenen `h1`-Selektor nicht ändern.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

Du solltest kein neues `style`-Tag hinzufügen. Füge die neuen CSS-Regeln zum vorhandenen `style`-Tag hinzu.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Du solltest einen neuen `h2`-Selektor hinzufügen.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

Du solltest einen neuen `p`-Selektor hinzufügen.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

Dein `h1`-Element sollte ein `text-align` von `center` haben.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

Dein `h2`-Element sollte ein `text-align` von `center` haben.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

Dein `p`-Element sollte ein `text-align` von `center` haben.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

