---
id: 5f3477ae34c1239cafe128be
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

You now have three type selectors with the same styling. You can add the same group of styles to many elements by creating a list of selectors. Each selector is separated with commas like this:

```css
selector1, selector2 {
  property: value;
}
```

Lösche die drei vorhandenen Typselektoren und ersetze sie durch eine Selektorliste, die den Text der `h1`-, `h2`-, und `p`-Elemente zentriert.

# --hints--

Du solltest für alle drei Elemente einen einzigen Typselektor wählen: `h1, h2, p`. Stelle sicher, dass du diese Reihenfolge benutzt.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

Du solltest nur einen Selektor in deinem `style`-Element haben.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

