---
id: 5f3477aefa51bfc29327200b
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

You have styled three elements by writing CSS inside the `style` tags. This works, but since there will be many more styles, it's best to put all the styles in a separate file and link to it.

A separate `styles.css` file has been created for you. You can switch between files with the tabs at the top of the editor.

Beginne damit, die von dir erstellten Formate in die Datei `styles.css` zu schreiben. Make sure you exclude the opening and closing `style` tags.

# --hints--

Deine `styles.css`-Datei sollte den `h1, h2, p`-Typselektor enthalten.

```js
assert.match(editableContents.replace(/[\s\n]*/g, ""), /(h1|h2|p),(h1|h2|p),(h1|h2|p){/);
```

Dein Selektor sollte die `text-align`-Eigenschaft auf `center` setzen.

```js
assert.match(editableContents, /text-align:\s*center;?/);
```

Dein Code sollte keine `<style>`-Tags enthalten.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
