---
id: 67c562286b29447da020d407
title: 建置一次性密碼產生器
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

在本實驗中，你將產生一組 6 位數的 OTP（一次性密碼）並顯示給使用者。OTP 將在 5 秒後過期，當使用者點擊 `"Generate New OTP"` 按鈕時，將會產生新的 OTP。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該使用 `useEffect` 掛鉤來管理倒數計時器。
2. 你的 `OTPGenerator` 組件應該傳回一個具有類別名稱 `container` 的 `div` 元素。
3. 具有類別 `container` 的 `div` 應包含以下元素：

    - 一個具有 ID `otp-title` 且文字為 `"OTP Generator"` 的 `h1` 元素。
    - 一個具有 ID `otp-display` 的 `h2` 元素，該元素要麼顯示訊息 `"Click 'Generate OTP' to get a code"`，要麼顯示已產生的 OTP（如果有的話）。
    - 一個具有 ID 為 `otp-timer` 且 `aria-live` 屬性設為有效值的 `p` 元素：
        - 起初是空的。
        - 在按鈕被點擊後顯示 `"Expires in: X seconds"`，其中 `X` 表示 OTP 過期前的剩餘時間。
        - 當倒數計時達到 `0` 時，顯示訊息 `"OTP expired. Click the button to generate a new OTP."`。
    - 一個 ID 為 `generate-otp-button` 的 `button` 元素，標籤為 `"Generate OTP"`。當點擊時，應該產生一個新的 OTP 並開始 5 秒倒數。
    - 當倒數計時處於作用中時，必須停用 `"Generate OTP"` 按鈕。

4. 你應該確保倒數計時器在達到 `0` 秒時自動停止，以防止不必要的更新。
5. 產生的 OTP 應該是 6 位數。

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

你不應該從程式碼中移除現有的 `const { useState, useEffect, useRef } = React;` 指派。

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

你應該呈現一個帶有類別 `.container` 的 `div`。

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

`.container` 應該包含一個 ID 為 `otp-title`、文字為 `"OTP Generator"` 的 `h1` 元素。

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

`.container` 應該包含一個具有 ID `otp-display` 的 `h2` 元素，該元素在 OTP 產生時顯示 OTP。 

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

最初，`.container` 內的 `h2` 應該顯示訊息 `"Click 'Generate OTP' to get a code"`。

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

`.container` 元素應包含具有 ID `otp-timer` 的 `p` 元素。

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

`p` 元素應該有一個 `aria-live` 屬性，設定為 `assertive` 或 `polite`。

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

最初，`p` 元素應該是空的。

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

`.container` 應該包含一個 ID 為 `generate-otp-button` 且文字為 `"Generate OTP"` 的 `button` 元素。 

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

當按鈕被點擊時，應該在具有 id `otp-display` 的 `h2` 元素中顯示一個新的 OTP。

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

產生的 OTP 應該是 6 位數。

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


當按鈕被點擊時，id 為 `otp-timer` 的 `p` 元素應顯示 5 秒倒數。

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

具有 id `otp-timer` 的 `p` 元素中的訊息應該每秒更新一次，以顯示剩餘時間。

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

當倒數計時器正在執行時，`"Generate OTP"` 按鈕應該被停用。

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

當倒數計時器達到 0 時，應該啟用 `"Generate OTP"` 按鈕。

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

當倒數計時器達到 `0` 時，你應該顯示訊息 `OTP expired. Click the button to generate a new OTP.`。

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

你應該匯出 `OTPGenerator` 組件。

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
