---
id: 67c562286b29447da020d407
title: 建置一次性密碼產生器
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

在這個實驗中，你將產生一組 6 位數的 OTP（一次性密碼）並顯示給使用者。OTP 會在 5 秒後過期，當使用者點擊 `"Generate OTP"` 按鈕時，會產生新的 OTP。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該使用 `useEffect` 掛鉤來管理倒數計時器。
2. 你的 `OTPGenerator` 組件應該回傳一個帶有類別名稱 `container` 的 `div` 元素。
3. 你的帶有類別 `container` 的 `div` 應包含以下元素：

    - 一個 ID 為 `otp-title`，文字為 `"OTP Generator"` 的 `h1` 元素。
    - 一個 ID 為 `otp-display` 的 `h2` 元素，該元素要顯示訊息 `"Click 'Generate OTP' to get a code"`，或在有產生 OTP 時顯示該 OTP。
    - 一個 ID 為 `otp-timer`，並且 `aria-live` 屬性設為有效值的 `p` 元素，該元素：
        - 初始為空。
        - 在按鈕被點擊後顯示 `"Expires in: X seconds"`，其中 `X` 代表 OTP 過期前剩餘的秒數。
        - 當倒數計時到 `0` 時，顯示訊息 `"OTP expired. Click the button to generate a new OTP."`。
    - 一個 ID 為 `generate-otp-button`，標示為 `"Generate OTP"` 的 `button` 元素。點擊時應產生新的 OTP 並開始 5 秒倒數。
    - 當倒數計時作用中時，`"Generate OTP"` 按鈕必須被停用。

4. 你應確保倒數計時器在達到 `0` 秒時自動停止，以避免不必要的更新。
5. 你產生的 OTP 應該是 6 位數長。

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

你不應該從程式碼中移除現有的 `const { useState, useEffect, useRef } = React;` 指派。

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

你應該呈現一個帶有類別 `.container` 的 `div`。

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

你的 `.container` 應包含一個 ID 為 `otp-title`，文字為 `"OTP Generator"` 的 `h1` 元素。

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

你的 `.container` 應包含一個 ID 為 `otp-display`，用來顯示產生的 OTP 的 `h2` 元素。

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

你 `.container` 裡的 `h2` 元素應該初始顯示訊息 `"Click 'Generate OTP' to get a code"`。

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

你的 `.container` 元素應包含一個 ID 為 `otp-timer` 的 `p` 元素。

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

你的 `p` 元素應該有一個 `aria-live` 屬性，設為 `assertive` 或 `polite`。

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

你的 `p` 元素初始應該是空的。

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

你的 `.container` 應包含一個 ID 為 `generate-otp-button`，文字為 `"Generate OTP"` 的 `button` 元素。

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

當按鈕被點擊時，ID 為 `otp-display` 的 `h2` 元素應顯示新的 OTP。

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

你產生的 OTP 應該是 6 位數長。

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


當按鈕被點擊時，ID 為 `otp-timer` 的 `p` 元素應顯示 5 秒倒數。

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

ID 為 `otp-timer` 的 `p` 元素應每秒更新訊息以顯示剩餘時間。

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

當倒數計時器正在執行時，`"Generate OTP"` 按鈕應該被停用。

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

當倒數計時器達到 0 時，`"Generate OTP"` 按鈕應該被啟用。

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

當倒數計時器達到 `0` 時，你應該顯示訊息 `OTP expired. Click the button to generate a new OTP.`。

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

你應該匯出 `OTPGenerator` 組件。

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
