---
id: 66a207974c806a19d6607073
title: Crie uma aplicação de inventário de livros
challengeType: 25
dashedName: build-a-book-inventory-app
demoType: onClick
---

# --description--


Atenda às histórias de usuário abaixo e faça todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter um elemento `h1` com o texto `Book Inventory`.
1. Você deve ter um elemento `table` com colunas nomeadas `Title`, `Author`, `Category`, `Status` e `Rate`.
1. A tabela deve ter pelo menos quatro linhas, a primeira para os cabeçalhos das colunas e o restante preenchido com informações do livro.
1. Cada linha de tabela dentro do corpo da tabela deve ter a classe `read`, `to-read` ou `in-progress`.
1. Elementos `td` da coluna `Status` devem conter um elemento `span` com a `class` de `status` envolvendo o texto `Read`, `To Read` ou `In Progress`, dependendo da classe daquela linha.
1. Elementos `td` da coluna `Rate` devem conter um elemento `span` com a `class` de `rate` envolvendo três elementos `span` vazios.
1. Elementos `.rate` colocados dentro de linhas `.read` devem ter uma classe adicional com o valor de `one`, `two` ou `three`, dependendo da classificação pessoal. Esse valor deve vir depois de `rate`.
1. Você deve criar três seletores de atributo para direcionar os elementos com as classes `read`, `to-read` e `in-progress`, além de definir sua propriedade `background-image` para usar um `linear-gradient` de sua escolha.
1. Você deve definir a propriedade `display` de cada elemento `span` como `inline-block`.
1. Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `to-read` e definir suas propriedades `border` e `background-image`.
1. Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `read` e definir suas propriedades `border` e `background-image`.
1. Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `in-progress` e definir suas propriedades `border` e `background-image`.
1. Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` e os elementos `span` com o valor da classe começando com `rate` e definir suas propriedades `height`, `width` e `padding`.
1. Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir suas propriedades `border`, `border-radius`, `margin`, `height`, `width` e `background-color`.
1. You should use an attribute selector to target the first descendant of `span` elements that have the word `one` as a part of their `class` value and set its `background-image` property to use a `linear-gradient`.
1. You should use an attribute selector to target the first two descendants of `span` elements that have the word `two` as a part of their `class` value and set their `background-image` property to use a `linear-gradient`.
1. You should use an attribute selector to target the three `span` elements that are descendants of `span` elements that have the word `three` as a part of their `class` value and set their `background-image` property to use a `linear-gradient`.

# --hints--

Você deve ter um elemento `h1` com o texto `Book Inventory`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Book Inventory');
```

Você deve ter apenas um elemento `h1`.

```js
assert.equal(document.querySelectorAll('h1')?.length, 1);
```

You should have a `table` element.

```js
assert.equal(document.querySelectorAll('table')?.length, 1);
```

O elemento `table` deve ter cinco colunas.

```js
assert.equal(document.querySelectorAll('th')?.length, 5);
```

A primeira coluna deve ter o texto `Title` como o cabeçalho.

```js
assert.equal(document.querySelectorAll('th')[0]?.innerText, 'Title');
```

A segunda coluna deve ter o texto `Author` como o cabeçalho.

```js
assert.equal(document.querySelectorAll('th')[1]?.innerText, 'Author');
```

A terceira coluna deve ter o texto `Category` como o cabeçalho.

```js
assert.equal(document.querySelectorAll('th')[2]?.innerText, 'Category');
```

A quarta coluna deve ter o texto `Status` como o cabeçalho.

```js
assert.equal(document.querySelectorAll('th')[3]?.innerText, 'Status');
```

A quinta coluna deve ter o texto `Rate` como o cabeçalho.

```js
assert.equal(document.querySelectorAll('th')[4]?.innerText, 'Rate');
```

A tabela deve ter pelo menos quatro linhas.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
```

Cada linha da tabela, exceto a linha do cabeçalho, deve ter a classe `read`, `to-read` ou `in-progress`.

```js
const rows = document.querySelectorAll('tr');
assert.isAtLeast(rows.length, 4);
for (let i = 1; i < rows.length; i++) {
  const classList = [...rows[i].classList];
  const currentProgress = classList[0]; 
  assert.oneOf(currentProgress, ['read', 'to-read', "in-progress"]); 
}
```

Elementos `td` da coluna `Status` devem conter um elemento `span`.

```js
const statusColumnData = document.querySelectorAll('td:nth-child(4)');
assert.isAbove(statusColumnData.length, 0);
for (let e of statusColumnData) {
  assert.equal(e?.children[0]?.tagName, 'SPAN');
}
```

Cada elemento `span` da coluna `Status` deve ter a classe de `status`.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
assert.isAbove(statusSpans.length, 0);
for (let e of statusSpans) {
    assert.isTrue(e.classList.contains('status'));
}
```

Cada elemento `.status` deve ter o texto `Read`, `To Read` ou `In Progress`, dependendo da classe de sua linha.

```js
const statusSpans = document.querySelectorAll('tr td:nth-child(4) :first-child');
const rows = Array.from(document.querySelectorAll('tr')).slice(1);
assert.isAbove(statusSpans.length, 0);
for (let i = 0; i < rows.length; i++) {
  switch (statusSpans[i]?.innerText) {
        case 'Read':
            assert.isTrue(rows[i].classList.contains('read'));
            break;
        case 'To Read':
            assert.isTrue(rows[i].classList.contains('to-read'));
            break;
        case 'In Progress':
            assert.isTrue(rows[i].classList.contains('in-progress'));
            break;
        default:
            assert.fail(); 
    }
}
```

Elementos `td` da coluna `Rate` devem conter um elemento `span`.

```js
const rateColumnData = document.querySelectorAll('tr td:last-child');
assert.isAbove(rateColumnData.length, 0);
for (let e of rateColumnData) {
    assert.equal(e.children[0]?.tagName, 'SPAN')
}
```

Cada elemento `span` que é filho direto de um elemento `td` da coluna `Rate` deve ter a classe de `rate` como a primeira classe.

```js
const rateSpans = document.querySelectorAll('tr td:last-child > span:first-child');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.classList[0], 'rate');
}
```

Cada elemento `.rate` deve conter três elementos `span` vazios.

```js
const rateSpans = document.getElementsByClassName('rate');
assert.isAbove(rateSpans.length, 0);
for (let e of rateSpans) {
    assert.equal(e.children.length, 3);
    for (let child of e.children) {
        assert.equal(child.tagName, 'SPAN');
        assert.equal(child.innerText.length, 0);
    }
}
```

Elementos `.rate` colocados dentro de linhas `.read` devem ter uma classe adicional após a classe `rate` com o valor de `one`, `two` ou `three`, dependendo da classificação pessoal.

```js
const readBooksRates = document.querySelectorAll('.read .rate');
assert.isAbove(readBooksRates.length, 0);
for (let e of readBooksRates) {
    assert.oneOf(e.classList[1], ['one', 'two', 'three']);
}
```

Você deve ter um seletor de atributo para direcionar as linhas que têm a classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"]'));
```

Você deve usar um seletor de atributo para direcionar as linhas que têm a classe `read` e definir sua propriedade `background-image` para um gradiente linear à sua escolha.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="read"]')?.backgroundImage, 'linear-gradient(');
```

Você deve ter um seletor de atributo para direcionar as linhas que têm a classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]'));
```

Você deve usar um seletor de atributo para direcionar as linhas que têm a classe `to-read` e definir sua propriedade `background-image` para um gradiente linear à sua escolha.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"]')?.backgroundImage, 'linear-gradient(');
```

Você deve ter um seletor de atributo para direcionar as linhas que têm a classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]'));
```

Você deve usar um seletor de atributo para direcionar as linhas que têm a classe `in-progress` e definir sua propriedade `background-image` para um gradiente linear à sua escolha.

```js
assert.include(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"]')?.backgroundImage, 'linear-gradient(');
```

Você deve definir a propriedade `display` de cada elemento `span` como `inline-block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('span')?.getPropVal('display'), 'inline-block');
```

Você deve ter um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `to-read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `to-read` e definir sua propriedade `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.border);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `to-read` e definir sua propriedade `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="to-read"] span[class="status"]')?.backgroundImage);
```

Você deve ter um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `read`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `read` e definir sua propriedade `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.border);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `read` e definir sua propriedade `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="read"] span[class="status"]')?.backgroundImage);
```

Você deve ter um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `in-progress`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `in-progress` e definir sua propriedade `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.border);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` que são descendentes de elementos `tr` com a classe `in-progress` e definir sua propriedade `background-image`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('tr[class="in-progress"] span[class="status"]')?.backgroundImage);
```

Você deve ter um seletor de atributo para direcionar os elementos `span` com a classe `status` e os elementos `span` com o valor da classe começando com `rate`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.exists(selector1 || selector2);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` e os elementos `span` com o valor da classe começando com `rate` e definir sua propriedade `height`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.height || selector2?.height);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` e os elementos `span` com o valor da classe começando com `rate` e definir sua propriedade `width`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.width || selector2?.width);
```

Você deve usar um seletor de atributo para direcionar os elementos `span` com a classe `status` e os elementos `span` com o valor da classe começando com `rate` e definir sua propriedade `padding`.

```js
const selector1 = new __helpers.CSSHelp(document).getStyle('span[class="status"], span[class^="rate"]');
const selector2 = new __helpers.CSSHelp(document).getStyle('span[class^="rate"], span[class="status"]');
assert.isNotEmpty(selector1?.padding || selector2?.padding);
```

Você deve ter um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `border`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `border-radius`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('border-radius'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `margin`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('margin'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `height`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('height'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('width'));
```

Você deve usar um seletor de atributo para direcionar os elementos `span` que são filhos diretos de elementos `span` com o valor da `class` começando com `rate` e definir sua propriedade `background-color`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('span[class^="rate"] > span')?.getPropVal('background-color'));
```

You should have an attribute selector to target the first descendant of `span` elements that have the word `one` as a part of their `class` value.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

You should use an attribute selector to target the first descendant of `span` elements that have the word `one` as a part of their `class` value and set its `background-image` property to use a `linear-gradient`.

```js
const selectors = [
  'span[class~="one"] :first-child',
  'span[class~="one"] :nth-child(1)',
  'span[class~="one"] :first-of-type',
  'span[class~="one"] span:first-child',
  'span[class~="one"] span:nth-child(1)',
  'span[class~="one"] span:first-of-type',
  'span[class~="one"] > :first-child',
  'span[class~="one"] > :nth-child(1)',
  'span[class~="one"] > :first-of-type',
  'span[class~="one"] > span:first-child',
  'span[class~="one"] > span:nth-child(1)',
  'span[class~="one"] > span:first-of-type'
]
assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image').includes('linear-gradient('));
```

You should have an attribute selector to target the first two descendants of `span` elements that have the word `two` as a part of their `class` value.

```js

const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]
assert.isNotNull(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

You should use an attribute selector to target the first two descendants of `span` elements that have the word `two` as a part of their `class` value and set their `background-image` property to use a `linear-gradient`.

```js
const selectors = [
  'span[class~="two"] :nth-child(1), span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :nth-child(1)',
  'span[class~="two"] :first-child, span[class~="two"] :nth-child(2)',
  'span[class~="two"] :nth-child(2), span[class~="two"] :first-child',
  'span[class~="two"] :nth-of-type(-n+2)',
  'span[class~="two"] :nth-child(-n+2)',
  'span[class~="two"] span:nth-child(1), span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:nth-child(1)',
  'span[class~="two"] span:first-child, span[class~="two"] span:nth-child(2)',
  'span[class~="two"] span:nth-child(2), span[class~="two"] span:first-child',
  'span[class~="two"] span:nth-of-type(-n+2)',
  'span[class~="two"] span:nth-child(-n+2)',
  'span[class~="two"] > :nth-child(1), span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :nth-child(1)',
  'span[class~="two"] > :first-child, span[class~="two"] > :nth-child(2)',
  'span[class~="two"] > :nth-child(2), span[class~="two"] > :first-child',
  'span[class~="two"] > :nth-of-type(-n+2)',
  'span[class~="two"] > :nth-child(-n+2)',
  'span[class~="two"] > span:nth-child(1), span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:nth-child(1)',
  'span[class~="two"] > span:first-child, span[class~="two"] > span:nth-child(2)',
  'span[class~="two"] > span:nth-child(2), span[class~="two"] > span:first-child',
  'span[class~="two"] > span:nth-of-type(-n+2)',
  'span[class~="two"] > span:nth-child(-n+2)'
  ]

assert.isTrue(new __helpers.CSSHelp(document).getStyleAny(selectors)?.backgroundImage.includes('linear-gradient(')) 
```

You should have an attribute selector to target the `span` elements that are descendants of `span` elements that have the word `three` as a part of their `class` value.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.exists(new __helpers.CSSHelp(document).getStyleAny(selectors));
```

You should use an attribute selector to target the `span` elements that are descendants of `span` elements that have the word `three` as a part of their `class` value and set their `background-image` property to use a `linear-gradient`.

```js
const selectors = [
  'span[class~="three"] span',
  'span[class~="three"] > span'
]
assert.include(new __helpers.CSSHelp(document).getStyleAny(selectors)?.getPropVal('background-image'), 'linear-gradient(');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Book Inventory</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main>
    <h1>Book Inventory</h1>
    <div class="table-container">
      <table>
        <thead>
          <tr>
            <th>Title</th>
            <th>Author</th>
            <th>Category</th>
            <th>Status</th>
            <th>Rate</th>
          </tr>
        </thead>
        <tbody>
          <tr class="read">
            <td>The Three Musketeers</td>
            <td>Alexandre Dumas</td>
            <td>Historical Novel</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate three">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="to-read">
            <td>The Plague</td>
            <td>Albert Camus</td>
            <td>Philosophical Novel</td>
            <td><span class="status">To Read</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>The Metamorphosis</td>
            <td>Franz Kafka</td>
            <td>Novella</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate one">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Dead Souls</td>
            <td>Nikolai Gogol</td>
            <td>Picaresque</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="in-progress">
            <td>Lord of the Flies</td>
            <td>William Golding</td>
            <td>Allegorical Novel</td>
            <td><span class="status">In Progress</span></td>
            <td>
              <span class="rate">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
          <tr class="read">
            <td>Do Androids Dream of Electric Sheep?</td>
            <td>Philip K. Dick</td>
            <td>Science Fiction</td>
            <td><span class="status">Read</span></td>
            <td>
              <span class="rate two">
                <span></span><span></span><span></span>
              </span>
            </td>
          </tr>
        </tbody>
      </table>
    </div>
  </main>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
    font-family: Arial, sans-serif;
}

.table-container {
  overflow-x: auto;
}

table {
    border-collapse: collapse;
    width: 100vw;
}

thead {
    background-image: linear-gradient(hsl(197, 92%, 77%, 0.60), hsl(197, 92%, 50%, 0.60));
}

th {
    padding: 0.4em;
}

td {
    text-align: center;
    padding: 0.3em;
}

tr[class="read"] {
    background-image: linear-gradient(hsl(120, 100%, 85%, 0.6), hsl(120, 100%, 65%, 0.6));
}

tr[class="to-read"] {
    background-image: linear-gradient(hsl(120, 20%, 95%, 0.6), hsl(120, 20%, 75%, 0.6));
}

tr[class="in-progress"] {
    background-image: linear-gradient(hsl(40, 100%, 85%, 0.60), hsl(40, 100%, 65%, 0.6));
}

span {
    display: inline-block;
}

span[class="status"] {
    border-radius: 20%/60%;
}

span[class="status"],
span[class^="rate"] {
    height: 1.85em;
    width: 7em;
    padding: 0.3em;
}

tr[class="to-read"] span[class="status"] {
    border: 0.1em solid hsl(5, 100%, 45%);
    background-image: linear-gradient(hsl(5, 100%, 75%), hsl(5, 100%, 50%));
}

tr[class="read"] span[class="status"] {
    border: 0.1em solid hsl(120, 75%, 45%);
    background-image: linear-gradient(hsl(120, 75%, 75%), hsl(120, 75%, 50%));
}

tr[class="in-progress"] span[class="status"] {
    border: 0.1em solid hsl(40, 90%, 40%);
    background-image: linear-gradient(hsl(40, 90%, 75%), hsl(40, 90%, 50%));
}

span[class^="rate"] > span {
    border: 0.1em solid hsl(0, 0%, 50%);
    border-radius: 50%;
    background-color: hsl(0, 15%, 95%);
    height: 100%;
    width: 20%;
    margin: 0.1em;
}

span[class~="three"] span {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="two"] :nth-child(1),
span[class~="two"] :nth-child(2) {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}

span[class~="one"] :first-child {
    background-image: linear-gradient(hsl(50, 100%, 70%), hsl(50, 100%, 50%));
}
```
