---
id: 66f17db06803d11a1bd19a20
title: بناء لعبة اختبار قصير
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ مصفوفة باسم `questions`.
1. يجب أن تحتوي مصفوفة `questions` على خمسة كائنات على الأقل، كل منها يحتوي على المفاتيح `category` و`question` و`choices` و`answer`.
1. يجب أن تكون قيمة المفتاح `category` سلسلة نصية تمثل فئة السؤال.
1. يجب أن تكون قيمة المفتاح `question` سلسلة نصية تمثل السؤال.
1. يجب أن تكون قيمة المفتاح `choices` مصفوفة تحتوي على ثلاث سلاسل نصية، وهي إجابات بديلة للسؤال.
1. يجب أن تكون قيمة المفتاح `answer` سلسلة نصية تمثل الإجابة الصحيحة للسؤال. كما يجب أن تكون قيمة `answer` مدرجة في مصفوفة `choices`.
1. يجب أن تكون لديك دالة باسم `getRandomQuestion` تأخذ مصفوفة من الأسئلة كمعلمة وتُرجع كائن سؤال عشوائي من المصفوفة.
1. يجب أن تكون لديك دالة باسم `getRandomComputerChoice` تأخذ مصفوفة الخيارات المتاحة كمعلمة، وتُرجع إجابة عشوائية للسؤال المختار.
1. يجب أن تكون لديك دالة باسم `getResults` تأخذ كائن السؤال كالمعلمة الأولى وخيار الكمبيوتر كالمعلمة الثانية. يجب أن تُرجع الدالة `The computer's choice is correct!` إذا كانت الإجابة صحيحة. وإلا، تُرجع `The computer's choice is wrong. The correct answer is: <correct-answer>`، حيث `<correct-answer>` هي قيمة الإجابة الصحيحة للسؤال المختار.

# --hints--

يجب أن تنشئ مصفوفة باسم `questions`.

```js
assert.isArray(questions);
```

يجب أن تحتوي مصفوفة `questions` على خمسة كائنات على الأقل، كل منها يحتوي على المفاتيح `category` و`question` و`choices` و`answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

يجب أن تكون قيمة المفتاح `category` سلسلة نصية تمثل فئة السؤال.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

يجب أن تكون قيمة المفتاح `question` سلسلة نصية تمثل السؤال.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

يجب أن تكون قيمة المفتاح `choices` مصفوفة تحتوي على ثلاث سلاسل نصية مختلفة عن بعضها.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

يجب أن تكون قيمة المفتاح `answer` سلسلة نصية.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

يجب أن تكون قيمة `answer` مدرجة في مصفوفة `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

يجب أن تكون لديك دالة باسم `getRandomQuestion` تأخذ مصفوفة من الأسئلة كمعلمة وتُرجع كائن سؤال عشوائي من المصفوفة.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

يجب أن تكون لديك دالة باسم `getRandomComputerChoice` تأخذ مصفوفة الخيارات المتاحة كمعلمة، وتُرجع إجابة عشوائية للسؤال المختار.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

يجب أن تكون لديك دالة باسم `getResults`.

```js
assert.isFunction(getResults);
```

يجب أن تأخذ دالة `getResults` كائن السؤال كالمعلمة الأولى وخيار الكمبيوتر كالمعلمة الثانية.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

إذا كان خيار الكمبيوتر يطابق الإجابة، يجب أن تُرجع دالة `getResults` النص `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

إذا لم يطابق خيار الكمبيوتر الإجابة، يجب أن تُرجع دالة `getResults` النص `The computer's choice is wrong. The correct answer is: <correct-answer>`، حيث `<correct-answer>` هي قيمة الإجابة الصحيحة للسؤال المختار.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

يجب أن تستخدم دالة `getResults` مقارنة تساوي دقيقة، وليس مطابقة جزئية.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
