---
id: 68e707257ea91c2e12d489fb
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você praticará como adicionar estilos personalizados aos botões de opção construindo um formulário para conferência entre pais e professores. O boilerplate HTML foi fornecido para você.

Comece adicionando um elemento `main` com uma classe chamada `container`.

# --hints--

Você deve ter apenas um elemento `main`.

```js
assert.exists(document.querySelector("main"));
```

Seu elemento `main` deve ter uma classe chamada `container`.

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
