---
id: 68e7e729f6d14f69a23786ef
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Em seguida, dentro do seu elemento `main`, adicione um elemento `h1` com o texto `Parent Teacher Conference Form` e as classes `title` e `center`. 

# --hints--

Você deve ter um elemento `h1`.

```js
assert.exists(document.querySelector("h1"));
```

Seu elemento `h1` deve conter o texto `Parent Teacher Conference Form`.

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

Seu elemento `h1` deve ter as classes `title` e `center`.

```js
assert.exists(document.querySelector("h1.title.center"));
```

Seu elemento `h1` deve estar dentro do elemento `main`.

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  --fcc-editable-region--
  <main class="container">
    
  </main>
  --fcc-editable-region--
</body>

</html>
```
