---
id: 68ea8c330ed07b60879f3d38
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Agora é hora de adicionar os elementos do formulário para coletar as informações da nota do aluno.

Comece adicionando outro elemento `label` com o texto `Student Grade: ` e o atributo `for` definido como `"grade"`.

Então, abaixo do seu elemento `label`, adicione um elemento `input` com o atributo `type` definido como `"number"`. Os atributos `name` e `id` devem ser definidos como `"grade"`. O atributo `placeholder` deve ser definido como `"E.g., 4"`. Por fim, seu `input` do tipo número deve ser obrigatório. 

# --hints--

Você deve ter um total de dois elementos `label` na página.

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

Seu segundo elemento `label` deve ter o texto `Student Grade: `. Não esqueça o espaço após os dois pontos.

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

Seu segundo elemento `label` deve ter um atributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

O atributo `for` do seu segundo elemento `label` deve ter o valor `"grade"`.

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

Você deve ter dois elementos `input` na página.

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

Seu segundo elemento `input` deve estar abaixo do seu segundo elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Seu segundo elemento `input` deve ter um atributo `type` definido como `"number"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

Seu segundo elemento `input` deve ter um atributo `name` definido como `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

Seu segundo elemento `input` deve ter um atributo `id` definido como `"grade"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

Seu segundo elemento `input` deve ter um atributo `placeholder` definido como `"E.g., 4"`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

Seu segundo elemento `input` deve ter um atributo `required`.

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
