---
id: 68eaa668d4d9fbe529a0e16b
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Agora, adicione um elemento `label` com o texto `Any specific concerns or topics you'd like to discuss?` e um atributo `for` definido como `notes`.

Abaixo deste `label`, adicione um elemento `textarea`. Defina seu `id` como `notes`, `name` como `notes`, `rows` como `4` e `cols` como `50`.

# --hints--

Seu último `fieldset` deve conter um elemento `label`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

O elemento `label` deve conter o texto `Any specific concerns or topics you'd like to discuss?`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

O atributo `for` do elemento `label` deve ser definido como `"notes"`.

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

Seu último `fieldset` deve conter um elemento `textarea`.

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

Seu elemento `textarea` deve estar abaixo do elemento `label`.

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

O elemento `textarea` deve ter um atributo `id` definido como `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

O elemento `textarea` deve ter um atributo `name` definido como `"notes"`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

O elemento `textarea` deve ter um atributo `rows` definido como `4`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

O elemento `textarea` deve ter um atributo `cols` definido como `50`.

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Information: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>

        --fcc-editable-region--
        
        --fcc-editable-region--

      </fieldset>
    </form>
  </main>
</body>

</html>
```
