---
id: 68ea81da9f687732368a5322
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Agora é hora de adicionar os elementos `form` e `input`, que representarão as informações do responsável e do aluno.

Comece adicionando um elemento `form` abaixo do elemento `p`.

# --hints--

Você deve ter um elemento `form`.

```js
assert.exists(document.querySelector(".container form"));
```

O elemento do formulário deve estar abaixo do elemento `p`.

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
