---
id: 68ea8666b8b473455f5b33ef
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

A primeira seção do formulário se concentrará nas informações do estudante, como o nome e a série.

Dentro do elemento `form`, adicione um elemento `fieldset`. Dentro desse elemento `fieldset`, adicione um elemento `legend` com o texto `Student Information`. 

Abaixo do elemento `legend`, adicione um elemento `label` com o texto `Full Name: ` e um atributo `for` com o valor `student-name`.


# --hints--

Você deve ter um elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset"));
```

O elemento `fieldset` deve estar dentro do elemento `form`.

```js
assert.exists(document.querySelector("form fieldset"));
```

Você deve ter um elemento `legend`.

```js
assert.exists(document.querySelector("legend"));
```

O elemento `legend` deve estar dentro do elemento `fieldset`.

```js
assert.exists(document.querySelector("fieldset legend"));
```

Seu elemento `legend` deve conter o texto `Student Information`.

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

Você deve ter um elemento `label`.

```js
assert.exists(document.querySelector("label"));
```

Seu elemento `label` deve conter o texto `Full Name: `. Não esqueça o espaço após os dois pontos.

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

Seu elemento `label` deve ter um atributo `for`.

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

O atributo `for` do seu elemento `label` deve ter o valor `student-name`.

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

Seu elemento `label` deve estar abaixo do elemento `legend`.

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    --fcc-editable-region--
    <form>

      
    </form>
    --fcc-editable-region--
  </main>
</body>

</html>
```
