---
id: 68ea8fdf2c9b167307e7f671
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Agora é hora de adicionar os elementos do formulário para coletar as informações dos pais.

Comece adicionando um elemento `label` com o texto `Parent/Guardian Name: ` e o atributo `for` definido como `"parent-name"`.

Em seguida, abaixo do seu elemento `label`, adicione um elemento `input` com o atributo `type` definido como `"text"`. Os atributos `name` e `id` devem ser definidos como `"parent-name"`. O atributo `placeholder` deve ser definido como `"E.g., Nancy Doe"`. Por fim, seu `input` deve ser obrigatório. 

# --hints--

Você deve ter um total de três elementos `label` na página.

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

Seu terceiro elemento `label` deve ter o texto `Parent/Guardian Name: `. Não esqueça o espaço após os dois pontos.

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

Seu terceiro elemento `label` deve ter um atributo `for`.

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

O atributo `for` do seu terceiro elemento `label` deve ter o valor `"parent-name"`.

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

Você deve ter três elementos `input` na página.

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

Seu terceiro elemento `input` deve estar abaixo do seu terceiro elemento `label`.

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

Seu terceiro elemento `input` deve ter um atributo `type` definido como `"text"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Seu terceiro elemento `input` deve ter um atributo `name` definido como `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

Seu terceiro elemento `input` deve ter um atributo `id` definido como `"parent-name"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

Seu terceiro elemento `input` deve ter um atributo `placeholder` definido como `"E.g., Nancy Doe"`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

Seu terceiro elemento `input` deve ter um atributo `required`.

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
