---
id: 68ea8a5db73b155708992687
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O próximo passo é adicionar o elemento `input` associado para as informações do estudante.

Comece adicionando um elemento `input` com um atributo `type` definido como `"text"`. Depois adicione um atributo `name` e `id`, ambos definidos como `student-name`. 

Em seguida, adicione um atributo `placeholder` definido como `E.g., Jane Doe`. E, por fim, adicione um atributo `required`. 

# --hints--

Você deve ter um elemento `input`.

```js
assert.exists(document.querySelector("fieldset input"));
```

Seu elemento `input` deve ter um atributo `type` definido como `"text"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

Seu elemento `input` deve ter um atributo `name` definido como `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

Seu elemento `input` deve ter um atributo `id` definido como `"student-name"`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

Seu elemento `input` deve ter um atributo `placeholder` definido como `E.g., Jane Doe`.

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

O elemento `input` deve ter o atributo `required`.

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      --fcc-editable-region--

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        
      </fieldset>

      --fcc-editable-region--

    </form>
  </main>
</body>

</html>
```
