---
id: 68ea9f01c51b66b937793611
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Em seguida, adicione outro label com uma classe `"contact-method"` e um atributo `for` definido como `"phone"`. O texto do label deve ser `Phone: `.

Abaixo do seu elemento `label`, adicione um botão de rádio com os atributos `id` e `value` definidos como `"phone"`. O atributo `name` deve ser definido como `"contact-method"` e a classe deve ser definida como `"contact-method-radio-btn"`.

# --hints--

Você deve ter dois elementos `label` no último `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

O elemento `label` deve ter uma classe `"contact-method"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

O atributo `for` do elemento `label` deve ser definido como `"phone"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

O elemento `label` deve conter o texto `Phone: `.

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

Você deve ter dois elementos `input` dentro do último `fieldset`.

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

Seu elemento `input` deve ter um atributo `type` definido como `radio`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

O elemento `input` deve ter um atributo `id` definido como `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

O elemento `input` deve ter um atributo `value` definido como `"phone"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

O elemento `input` deve ter um atributo `name` definido como `"contact-method"`.

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

O elemento `input` deve ter uma classe `"contact-method-radio-btn"`.

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        --fcc-editable-region--
        
        --fcc-editable-region--

      </fieldset>
    </form>
  </main>
</body>

</html>
```
