---
id: 68ea7c41f2555e16174e7cfa
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Em seguida, adicione um elemento de parágrafo com o texto `Please fill out the form below to help schedule your parent-teacher conference.`. Seu elemento de parágrafo também deve ter as classes `description` e `center`.

# --hints--

Você deve ter um elemento `p`.

```js
assert.exists(document.querySelector("p"));
```

Seu elemento `p` deve conter o texto `Please fill out the form below to help schedule your parent-teacher conference.`.

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

Seu elemento `p` deve ter as classes `description` e `center`.

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
