---
id: 616968c2c94c8071b349c146
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Ndani ya `body`, panga ndani kipengele cha `h1` chenye maandishi ya `CSS Color Markers`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `h1`.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Unapaswa kuwa na lebo ya ufunguzi ya `h1`.

```js
assert.match(code, /<h1\s*>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `h1`.

```js
assert.match(code, /<\/h1\s*>/i);
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele cha `body`.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `CSS Color Markers`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Kumbuka, herufi kubwa na ndogo pamoja na tahajia ni muhimu kwa maandishi ya `h1`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
