---
id: 616971b3cd990186b66c99a1
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa baada ya faili yako ya CSS ya nje kuandaliwa, unaweza kuanza kupamba ukurasa.

Kama ukumbusho, hivi ndivyo unavyolenga kipengele cha aya na kuipangilia kulia:

```css
p {
  text-align: right;
}
```

Tengeneza sheria mpya ya CSS inayolenga kipengele cha `h1`, na weka sifa yake ya `text-align` kuwa `center`.

# --hints--

Unapaswa kutumia kichaguzi cha `h1`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Sheria yako ya CSS ya `h1` inapaswa kuwa na sifa ya `text-align` iliyowekwa kuwa `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
