---
id: 61696ef7ac756c829f9e4048
title: Hatua ya 2
challengeType: 0
dashedName: step-2
---

# --description--

Katika mradi huu utatumia faili ya CSS ya nje kupamba ukurasa. Faili ya `styles.css` tayari imetengenezwa kwa ajili yako. Lakini kabla hujaweza kuitumia, utahitaji kuunganisha na ukurasa.

Panga ndani kipengele cha `link` ndani ya kipengele cha `head`. Kipe sifa ya kipengele cha `rel` yenye thamani `"stylesheet"` na sifa ya kipengele cha `href` yenye thamani `"styles.css"`.

# --hints--

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa kipengele kisicho na maudhui, hakipaswi kuwa na lebo ya kufunga `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `rel` yenye thamani `"stylesheet"`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `href` yenye thamani `"styles.css"`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
