---
id: 616d4a84b756d9c4b8255093
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Ni wakati wa kuongeza rangi kwenye alama. Kumbuka kwamba mojawapo ya njia za kuongeza rangi kwenye kipengele ni kutumia <dfn>neno msingi la rangi</dfn> kama `black`, `cyan`, au `yellow`.

Kama ukumbusho, hivi ndivyo unavyolenga darasa la vitu `freecodecamp`:

```css
.freecodecamp {
  
}
```

Tengeneza sheria mpya ya CSS inayolenga darasa la vitu `marker`, na weka sifa yake ya `background-color` kuwa `red`.

**Kumbuka:** Hutaona mabadiliko yoyote baada ya kuongeza CSS.

# --hints--

Unapaswa kutengeneza kichaguzi cha darasa la vitu ili kulenga darasa la vitu `marker`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Sheria yako ya CSS ya `.marker` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
