---
id: 616d525007b8c5d8b3308b1c
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Alama yako itakuwa nzuri zaidi ikiwa itakuwa katikati ya ukurasa. Njia rahisi ya kufanya hivyo ni kwa kutumia `margin` <dfn>sifa ya ufupisho</dfn>.

Unaweza kuweka eneo la ukingo la vipengele kando kando kwa sifa kama `margin-top` na `margin-left`, sifa ya ufupisho ya `margin` inafanya iwe rahisi kuweka maeneo mengi ya ukingo kwa wakati mmoja.

Ili kuweka alama yako katikati ya ukurasa, weka sifa yake ya `margin` kuwa `auto`. Hii itaweka `margin-top`, `margin-right`, `margin-bottom`, na `margin-left` zote kuwa `auto`.

# --hints--

Kanuni yako ya `.marker` CSS inapaswa kuwa na sifa ya `margin` iliyowekwa kuwa `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
