---
id: 616d595270d902f0e7086e18
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Wakati una vipengele vitatu tofauti vya alama `div`, vinaonekana kama mstatili mkubwa mmoja. Unapaswa kuongeza nafasi kati yao ili iwe rahisi kuona kila kipengele.

Wakati sifa ya `margin` yenye njia fupi ina thamani mbili, huweka `margin-top` na `margin-bottom` kwa thamani ya kwanza, na `margin-left` na `margin-right` kwa thamani ya pili.

Katika sheria yako ya `.marker` ya CSS, weka sifa ya `margin` kuwa `10px auto`.

# --hints--

Sheria yako ya `.marker` ya CSS inapaswa kuwa na sifa ya `margin` iliyowekwa kuwa `10px auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
