---
id: 616d3a67ccf800ad94ec89ae
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sasa utaongeza baadhi ya vipengele ambavyo hatimaye utavitengeneza kuwa alama za rangi.

Kwanza, ndani ya kipengele cha `body`, ongeza kipengele cha `div` chenye sifa ya kipengele ya `class` iliyowekwa kuwa `container`. Hakikisha kipengele cha `div` kimewekwa chini ya kipengele cha `h1`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `div`.

```js
assert.match(code, /<div.*>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `div`.

```js
assert.match(code, /<\/div\s*>/i);
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
assert.strictEqual(document.querySelector('div')?.parentElement?.localName, 'body');
```

Kipengele chako cha `div` kinapaswa kuwa baada ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Kipengele chako cha `div` kinapaswa kuwa na sifa ya kipengele ya `class` iliyowekwa kuwa `container`.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css
h1 {
  text-align: center;
}
```
