---
id: 617b92b9de349513466f6156
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Wakati alama za rangi nyekundu na buluu zinaonekana sawa, ile ya kijani ni nyepesi zaidi kuliko ilivyokuwa awali. Hii ni kwa sababu neno msingi la rangi `green` ni kivuli giza zaidi, na kiko karibu nusu kati ya rangi nyeusi na thamani kubwa zaidi ya kijani.

Katika sheria ya CSS ya `.two`, weka thamani ya kijani katika kitendakazi cha `rgb` kuwa `127` ili kupunguza nguvu yake.

# --hints--

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 127, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
