---
id: 617b97abd9f3f61d1590b815
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Katika mfano wa rangi wa kuongeza RGB, <dfn>rangi msingi</dfn> ni rangi ambazo, zinapochanganywa, huunda nyeupe safi. Lakini ili hili lifanikishe, kila rangi inahitaji kuwa katika kiwango chake cha juu kabisa cha nguvu.

Kabla ya kuchanganya rangi, weka alama yako ya kijani kurudi kwenye kijani safi. Kwa kitendakazi cha `rgb` katika sheria ya `.two` ya CSS, weka kijani kurudi kwenye thamani ya juu kabisa ya `255`.

# --hints--

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 127, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
