---
id: 617b8e0d93a8d10d9a90e720
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Tafadhali angalia kuwa `background-color` ya alama yako bado ni nyekundu. Hii ni kwa sababu uliweka thamani nyekundu ya kitendakazi cha `rgb` kuwa thamani kubwa zaidi ya `255`, au asilimia 100 nyekundu, na kuweka thamani za kijani na buluu zote kuwa `0`.

Sasa tumia kitendakazi cha `rgb` kuweka rangi zingine.

Katika sheria ya CSS ya `.two`, tumia kitendakazi cha `rgb` kuweka `background-color` kwa thamani kubwa zaidi ya kijani, na `0` kwa thamani zingine. Na katika sheria ya CSS ya `.three`, tumia kitendakazi cha `rgb` kuweka `background-color` kwa thamani kubwa zaidi ya buluu, na `0` kwa thamani zingine.

# --hints--

Sheria yako ya CSS ya `.two` haipaswi kutumia neno msingi la rangi la `green` kuweka sifa ya `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Sheria yako ya CSS ya `.two` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 255, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Sheria yako ya CSS ya `.three` haipaswi kutumia neno msingi la rangi la `blue` kuweka sifa ya `background-color`.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Sheria yako ya CSS ya `.three` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 255)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
