---
id: 619b74fa777a2b2473c94f82
title: Hatua ya 40
challengeType: 0
dashedName: step-40
---

# --description--

Update the CSS class selector `.two` so it targets the new `green` class. And update the `.three` class selector so it targets the new `blue` class.

# --hints--

Msimbo wako haufai tena kuwa na kichaguzi cha class ya `.two`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Unapaswa kutumia kichaguzi cha class kulenga class ya `green`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.green'));
```

Sheria yako ya CSS ya `.green` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 0, 0)');
```

Msimbo wako haufai tena kuwa na kichaguzi cha class ya `.three`.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Unapaswa kutumia kichaguzi cha class kulenga class ya `blue`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.blue'));
```

Sheria yako ya CSS ya `.blue` inapaswa kuwa na sifa ya `background-color` iliyowekwa kuwa `rgb(0, 0, 0)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

```
