---
id: 61a49d15bdbb5e57cc6fd280
title: Hatua 49
challengeType: 0
dashedName: step-49
---

# --description--

Vituo vya rangi hukuruhusu kusahihisha mahali ambapo rangi zimewekwa kando ya mstari wa tofauti ya rangi. Vinatumika katika kitendakazi cha `linear-gradient` kubainisha mahali ambapo rangi inamalizika na mabadiliko ya rangi inayofuata huanza. Iwe asilimia au thamani ya urefu inaweza kutumika kufafanua nafasi ya kituo cha rangi.

Kwa mfano, katika tofauti ya rangi nyekundu-nyeusi hii, mabadiliko kutoka nyekundu kwenda nyeusi hufanyika katika sehemu ya 90% kando ya mstari wa tofauti ya rangi, hivyo nyekundu inachukua sehemu kubwa ya nafasi iliyopo:

```css
linear-gradient(90deg, red 90%, black);
```

Katika kitendakazi cha `linear-gradient`, ongeza kituo cha rangi cha `75%` baada ya hoja ya rangi nyekundu ya kwanza. Usiongeze vituo vya rangi kwa hoja zingine za rangi.

# --hints--

Kanuni yako ya `.red` CSS inapaswa kuwa na sifa ya `background` iliyowekwa kuwa `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
