---
id: 61b093219e7fc020b43b1bb4
title: Step 69
challengeType: 0
dashedName: step-69
---

# --description--

Njia nyingine ya kuweka uwazi kwa kipengele ni kwa kutumia <dfn>alpha channel</dfn>. Kama sifa ya `opacity`, alpha channel inadhibiti jinsi rangi inavyoweza kuonekana wazi au kufunikwa.

Umeweka uwazi wa sleeve kwa kutumia rangi iliyopewa jina na sifa ya `opacity`, lakini unaweza kuongeza alpha channel kwa sifa nyingine za rangi za CSS.

Ndani ya sheria ya `.sleeve`, ondoa sifa na thamani ya `opacity`.

# --hints--

Sheria yako ya `.sleeve` ya CSS haipaswi kuwa na sifa na thamani ya `opacity`.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
--fcc-editable-region--
  opacity: 0.5;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
