---
id: 61b0936d9e7fc020b43b1bb8
title: Hatua 73
challengeType: 0
dashedName: step-73
---

# --description--

Inaonekana kama mkono wako umeondoka, lakini usijali — bado uko hapo. Kilichotokea ni kwamba kofia yako mpya `div` inachukua upana mzima wa alama, na inasukuma mkono chini hadi kwenye mstari unaofuata.

Hii ni kwa sababu sifa ya `display` ya kawaida kwa vipengele vya `div` ni `block`. Hivyo wakati vipengele viwili vya `block` viko karibu, vinaweka kama vitu halisi vinavyopangwa kwa safu. Kwa mfano, vipengele vyako vya alama vyote vimewekwa juu ya kila mmoja.

Ili kuweka vipengele viwili vya `div` kwenye mstari mmoja, weka sifa zao za `display` kuwa `inline-block`.

Tengeneza sheria mpya ili kulenga darasa la vitu la `cap` na `sleeve`, na weka `display` kuwa `inline-block`.

# --hints--

Unapaswa kutumia kichaguzi cha darasa la vitu kulenga darasa la `cap` na `sleeve`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve') || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap'));
```

Sheria yako ya CSS inapaswa kuwa na sifa ya `display` iliyowekwa kuwa `inline-block`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('.cap, .sleeve')?.display === 'inline-block' || new __helpers.CSSHelp(document).getStyle('.sleeve, .cap')?.display === 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
}

--fcc-editable-region--

--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
