---
id: 61b0a44a6b865738ba49b9fb
title: Step 81
challengeType: 0
dashedName: step-81
---

# --description--

Jambo la mwisho utakalofanya ni kuongeza kivuli kidogo kwa kila alama ili ziweze kuonekana halisi zaidi.

Sifa ya `box-shadow` inakuwezesha kutumia kivuli kimoja au zaidi kuzunguka kipengee. Hapa kuna sintaksia ya msingi:

```css
box-shadow: offsetX offsetY color;
```

Hivi ndivyo thamani za `offsetX` na `offsetY` zinavyofanya kazi:

* zote `offsetX` na `offsetY` zinakubali thamani za nambari katika `px` na vitengo vingine vya CSS
* thamani chanya ya `offsetX` husogeza kivuli kulia na thamani hasi husogeza kushoto
* thamani chanya ya `offsetY` husogeza kivuli chini na thamani hasi husogeza juu
* ikiwa unataka thamani ya sifuri (`0`) kwa yoyote au zote `offsetX` na `offsetY`, huna haja ya kuongeza kitengo. Kivinjari chochote huelewa kuwa sifuri inamaanisha hakuna mabadiliko.

Urefu na upana wa kivuli huamuliwa na urefu na upana wa kipengee kinachotumiwa. Pia unaweza kutumia thamani ya hiari ya `spreadRadius` ili kueneza wigo wa kivuli. Zaidi kuhusu hilo baadaye.

Anza kwa kuongeza kivuli rahisi kwa alama nyekundu.

Katika sheria ya CSS ya `.red`, ongeza sifa ya `box-shadow` yenye thamani `5px` kwa `offsetX`, `5px` kwa `offsetY`, na `red` kwa `color`.

# --hints--

Sheria yako ya CSS ya `.red` inapaswa kuwa na sifa ya ufupisho ya `box-shadow` na thamani `5px 5px red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red 5px 5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
