---
id: 61b306305810f1c9040ce5a2
title: Step 82
challengeType: 0
dashedName: step-82
---

# --description--

Kama unavyoona, umeongeza kivuli rahisi chepesi cha rangi nyekundu karibu na alama yako kilicho na pikseli 5 upande wa kulia, na pikseli 5 chini.

Lakini vipi kama ungependa kuweka kivuli chako upande wa pili? Unaweza kufanya hivyo kwa kutumia thamani hasi kwa `offsetX` na `offsetY`.

Sahihisha thamani za sifa ya `box-shadow`, na weka `offsetX` kuwa `-5px`, na `offsetY` kuwa `-5px`.

# --hints--

Kanuni yako ya `.red` ya CSS inapaswa kuwa na sifa ya ufupisho ya `box-shadow` na thamani `-5px -5px red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.boxShadow, 'red -5px -5px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
      <div class="marker blue">
        <div class="cap"></div>
        <div class="sleeve"></div>
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.cap {
  width: 60px;
  height: 25px;
}

.sleeve {
  width: 110px;
  height: 25px;
  background-color: rgba(255, 255, 255, 0.5);
  border-left: 10px double rgba(0, 0, 0, 0.75);
}

.cap, .sleeve {
  display: inline-block;
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
--fcc-editable-region--
  box-shadow: 5px 5px red;
--fcc-editable-region--
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
