---
id: 669aff9f5488f1bea056416d
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви будете працювати з семантичним HTML, створивши сторінку блогу, присвячену Віскерсу.

Щоб розпочати проєкт, додайте `<!DOCTYPE html>` та елемент `html` з атрибутом `lang` зі значенням `en`.

Пригадайте, що у попередньому модулі ви дізнались, як створити базовий шаблонний код на HTML.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

Ви повинні мати `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Ви повинні мати початковий тег `html` зі значенням англійської мови.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Ви повинні мати кінцевий тег `html`.

```js
assert.match(code, /<\/html>/i);
```

`DOCTYPE` має бути перед елементом `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
