---
id: 669fc7e141e4703748c558bf
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте елемент `head` до елемента `html`.

# --hints--

Ви повинні мати початковий тег `head`.

```js
assert.match(code, /<head>/i);
```

Ви повинні мати кінцевий тег `head`.

```js
assert.match(code, /<\/head>/i);
```

Початковий тег `head` має бути перед кінцевим тегом `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
