---
id: 669fd003cb89ee3c2402e041
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

В інформації для ознайомлення потрібно показати фотографію Віскерса з підписом.

Для початку додайте елемент `figure` під елементом `h1`.

Додайте елемент `img` до елемента `figure`.

The `src` attribute of the `img` should have a value of `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` and the `alt` text should have a value of `"a cat in the garden"`.

Below your `img` element, add a `figcaption` with the text `Mr. Whiskers in the Garden`.

# --hints--

Елемент `figure` повинен мати початковий тег.

```js
assert.match(code, /<figure\>/);
```

Елемент `figure` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figure\>/);
```

Елемент `figure` має бути в межах елемента `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Ви повинні мати елемент `img`.

```js
assert.exists(document.querySelector("img"));
```

Елемент `img` має бути в межах елемента `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Елемент `img` повинен мати атрибут `src` зі значенням `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Your `img` element should have an `alt` attribute set to `"a cat in the garden"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat in the garden');
```

Елемент `figcaption` повинен мати початковий тег.

```js
assert.match(code, /<figcaption\>/);
```

Елемент `figcaption` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figcaption\>/);
```

Елемент `figcaption` має бути в межах елемента `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Your `figcaption` element should have a the text of `Mr. Whiskers in the Garden`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText, "Mr. Whiskers in the Garden")
```

Елемент `figcaption` має бути під елементом `img`. Вони знаходяться в неправильному порядку.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```
