---
id: 669fcb5106eeb13ab7afcd09
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

The first section you will build out is the page <dfn>header</dfn>.

The `header` element is used to represent introductory content like page navigation and other introductory information.

Ось приклад використання елемента `header`:

```html
<header>
  <h1>Main Page Title Goes Here</h1>
  <img src="example-logo.png" alt="Example logo" />
</header>
```

Додайте елемент `header` до елемента `body`.

# --hints--

Ви повинні мати початковий тег `header`.

```js
assert.match(code, /<header>/i);
```

Ви повинні мати кінцевий тег `header`.

```js
assert.match(code, /<\/header>/i);
```

Your `header` element should be inside the `body` element.

```js
assert.match(code, /<\/header>[.\n\s]*<\/body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  --fcc-editable-region--
  <body>

  </body>
  --fcc-editable-region--
</html>
```
