---
id: 669fd003cb89ee3c2402e041
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

В інформації для ознайомлення потрібно показати фотографію Віскерса з підписом.

Для початку додайте елемент `figure` під елементом `h1`.

Додайте елемент `img` до елемента `figure`.

Елемент `img` повинен мати атрибути `src` зі значенням `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"` та `alt` зі значенням `"a cat peacefully sleeping"`.

Додайте елемент `figcaption` з текстом `Mr. Whiskers Sleeping` під елементом `img`.

# --hints--

Елемент `figure` повинен мати початковий тег.

```js
assert.match(code, /<figure\>/);
```

Елемент `figure` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figure\>/);
```

Елемент `figure` має бути в межах елемента `header`.

```js
const figureElement = document.querySelector('figure');
assert.equal(figureElement?.parentElement.tagName, 'HEADER');
```

Ви повинні мати елемент `img`.

```js
assert.exists(document.querySelector("img"));
```

Елемент `img` має бути в межах елемента `figure`.

```js
const imgElement = document.querySelector('img');
assert.strictEqual(imgElement?.parentElement.tagName, 'FIGURE');
```

Елемент `img` повинен мати атрибут `src` зі значенням `"https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"`.

```js
const imgElement = document.querySelector('img');
const src = imgElement?.getAttribute('src');

assert.strictEqual(src, 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Елемент `img` повинен мати атрибут `alt` зі значенням `"a cat peacefully sleeping"`.

```js
const imgElement = document.querySelector('img');
const alt = imgElement?.getAttribute('alt');

assert.strictEqual(alt, 'a cat peacefully sleeping');
```

Елемент `figcaption` повинен мати початковий тег.

```js
assert.match(code, /<figcaption\>/);
```

Елемент `figcaption` повинен мати кінцевий тег.

```js
assert.match(code, /<\/figcaption\>/);
```

Елемент `figcaption` має бути в межах елемента `figure`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.parentElement.tagName, 'FIGURE');
```

Елемент `figcaption` повинен містити текст `Mr. Whiskers Sleeping`.

```js
const figcaptionElement = document.querySelector('figcaption');
assert.strictEqual(figcaptionElement?.innerText.trim(), "Mr. Whiskers Sleeping")
```

Елемент `figcaption` має бути під елементом `img`. Вони знаходяться в неправильному порядку.

```js
const figureElement = document.querySelector('figure');
assert.strictEqual(figureElement?.lastElementChild?.tagName, "FIGCAPTION")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
--fcc-editable-region--
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>

    </header>
  </body>
--fcc-editable-region--
</html>
```
