---
id: 669fde2081f65141ad703fe4
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

The first section on the page will be the about section. The section will introduce Mr. Whiskers and give users an idea of what this blog is about.

Додайте елемент `section` з атрибутом `id` зі значенням `"about"` до елемента `main`.

Додайте елемент `h2` з текстом `About` до елемента `section`.

# --hints--

Елемент `section` повинен мати початковий тег.

```js
assert.match(code, /<section.*>/);
```

Елемент `section` повинен мати кінцевий тег.

```js
assert.match(code, /<\/section\>/);
```

Your `section` element should have an `id` set to `"about"`.

```js
const sectionElement = document.querySelector('section');
assert.strictEqual(sectionElement.getAttribute("id"), "about");
```

Елемент `section` має бути в межах елемента `main`.

```js
const sectionElement = document.getElementById('about');
assert.strictEqual(sectionElement.parentElement.tagName, 'MAIN');
```

Елемент `h2` повинен мати початковий тег.

```js
assert.match(code, /<h2\>/);
```

Елемент `h2` повинен мати кінцевий тег.

```js
assert.match(code, /<\/h2\>/);
```

Елемент `h2` повинен містити текст `About`. Double check your spelling.

```js
const h2Element = document.querySelector('body h2');
assert.strictEqual(h2Element.innerText, "About");
```

Елемент `h2` має бути в межах елемента `section`.

```js
const h2Element = document.querySelector('#about h2');
assert.strictEqual(h2Element.parentElement.tagName, 'SECTION');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
--fcc-editable-region--
    <main>

    </main>
--fcc-editable-region--
  </body>
</html>
```
