---
id: 669fd6fd12918e3de87854d4
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Користувачам блогу потрібна можливість переходити до різних розділів на сторінці.

Елемент `nav` використовують, щоб надати навігаційні посилання до інших розділів документа або вебсайту. Елемент `nav` часто використовують для меню або змісту.

Ось приклад використання елемента `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Під елементом `figure` додайте елемент `nav`, який містить елемент `ul`.

Додайте три елементи `li` до елемента `ul`.

# --hints--

Елемент `nav` повинен мати початковий тег.

```js
assert.match(code, /<nav\>/);
```

Елемент `nav` повинен мати кінцевий тег.

```js
assert.match(code, /<\/nav\>/);
```

Елемент `nav` має бути в межах елемента `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

Елемент `ul` повинен мати початковий тег.

```js
assert.match(code, /<ul\>/);
```

Елемент `ul` повинен мати кінцевий тег.

```js
assert.match(code, /<\/ul\>/);
```

Елемент `ul` має бути в межах елемента `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

Елемент `ul` повинен містити три елементи `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
