---
id: 669fcb06c3034a39f5431a38
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

To prepare creating some actual content, add a `body` element below the `head` element.

# --hints--

Ви повинні мати початковий тег `<body>`.

```js
assert.match(code, /<body>/i);
```

Ви повинні мати кінцевий тег `</body>`.

```js
assert.match(code, /<\/body>/i);
```

Не змінюйте елемент `head`. Переконайтеся, що не видалили кінцевий тег.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Елемент `body` повинен йти після елемента `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>

  --fcc-editable-region--
</html>
```
