---
id: 669fdfc9a5e3da42d2376609
title: Крок 13
challengeType: 0
dashedName: step-13
---

# --description--

Ви додали розділ «Про нас», тож спробуйте натиснути на посилання `About` і побачите, що сторінка перейде до цього розділу.

Наступний розділ сторінки міститиме список дописів, які стосуються Віскерса.

Додайте ще один елемент `section` з `id` зі значенням `"posts"`.

Додайте елемент `h2` з текстом `Posts` до елемента `section`.

# --hints--

Код повинен містити початковий тег `section`.

```js
assert.match(code, /<section.*>/i);
```

Код повинен містити кінцевий тег `section`.

```js
assert.match(code, /<\/section>/i);
```

Елемент `section` повинен мати `id` зі значенням `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

Елемент `section` має бути в межах елемента `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement?.parentElement.tagName, 'MAIN');
```

Елемент `h2` має бути в межах елемента `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.parentElement.tagName, 'SECTION');
```

Елемент `h2` повинен містити текст `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element?.innerText, 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
  </body>
</html>
```
