---
id: 66a7e72adf226c02626715a3
title: Крок 12
challengeType: 0
dashedName: step-12
---

# --description--

Під елементом `h2` додайте абзац з текстом `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.`.

Під цим абзацом додайте ще один абзац з текстом `His playful nature and boundless energy keeps me on my toes. I love him so much.`

# --hints--

Перший абзац повинен мати текст `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` Двічі перевірте написання.

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText, text);
```

Перший абзац має бути вкладеним в елементі `section`.

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

Другий абзац має бути в межах елемента `section`.

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

Другий абзац повинен мати текст `His playful nature and boundless energy keeps me on my toes. I love him so much.` Двічі перевірте написання.

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText, text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      --fcc-editable-region--
      <section id="about">
        <h2>About</h2>

      </section>
      --fcc-editable-region--
    </main>
  </body>
</html>
```
