---
id: 66b36358ed4f261d64840c24
title: بناء جدول التوفر
challengeType: 25
dashedName: build-an-availability-table
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي جدولك على ثلاثة أعمدة وخمسة صفوف على الأقل، بما في ذلك صفوف الرؤوس.
1. يجب أن يحتوي الصف الأول على خلايا رؤوس جدول تتضمن أيام الأسبوع.
1. يجب أن يحتوي العمود الأول على خلايا رؤوس جدول تحمل `class` باسم `time` وتتضمن قيم الوقت كنص لها.
1. يجب أن تحمل جميع خلايا البيانات `class` باسم `available-#`، حيث `#` رقم من `0` إلى `5` يحدد عدد الأشخاص المتاحين في ذلك الوقت.
1. يجب أن يحتوي جدولك على عنصر `caption` يحمل النص `Weekly Meeting Availability Schedule`.
1. يجب أن تحمل جميع خلايا رؤوس الجدول في الصف الأول (أيام الأسبوع) خاصية `scope` بالقيمة `col`.
1. يجب أن تحمل جميع خلايا رؤوس الجدول في العمود الأول (خلايا الوقت) خاصية `scope` بالقيمة `row`.
1. يجب أن تحمل جميع خلايا البيانات التي تحمل صنف التوفر خاصية `aria-label` تصف مستوى التوفر (مثلًا `0 available` لخلايا `available-0`، و `1 available` لخلايا `available-1`، وهكذا، و `5+ available` لخلايا `available-5`).
1. في محدد `:root`، يجب تعريف ستة متغيرات CSS بأسماء `--color#`، حيث `#` رقم من `0` إلى `5`، مع إسناد قيمة لون لكل منها. استخدم هذه المتغيرات لتعيين لون خلفية عناصر `.available-#` المقابلة.
1. في محدد `:root`، يجب تعريف متغيري CSS بالاسمين `--solid-border` و `--dashed-border`. استخدم هذين المتغيرين لتنسيق الحدود السفلية لخلايا البيانات، بالتناوب بين الحدود الصلبة والمتقطعة حسب الصف. امنح الصفوف إما الصنف `sharp` أو `half` لتنسيقها.
1. يجب أن يكون لديك عنصر `div` يحمل `id` بقيمة `legend`. ويجب أن يحتوي على عنصر `span` يحمل النص `Availability` وعنصر `div` يحمل `id` بقيمة `legend-gradient`.
1. يجب أن تمنح عنصر `#legend-gradient` تدرجًا لونيًا خطيًا ينتقل بين جميع الألوان من `--color0` إلى `--color5` بخطوط حادة.

**ملاحظة:** تأكد من ربط ورقة الأنماط الخاصة بك في ملف HTML لتطبيق CSS.

# --hints--

يجب أن يكون لديك عنصر `table`.

```js
assert.isNotNull(document.querySelector('table'));
```

يجب أن يكون لديك عنصر `table` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('table'), 1);
```

يجب أن يحتوي جدولك على `5` صفوف على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('tr')?.length, 5);
```

يجب أن يحتوي الجدول على ما لا يقل عن `3` أعمدة.

```js
const rows = [...document.querySelectorAll('tr')];
assert.isNotEmpty(rows);
rows.forEach((row) => {
    assert.isAtLeast(row.children.length, 3);
})
```

يجب أن يحتوي الصف الأول في جدولك على ثلاثة عناصر `th` على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('tr:first-child>th').length, 3)
```

يجب ألا يحتوي الصف الأول في جدولك على عناصر `td`.

```js
const firstRow = document.querySelector('tr');
assert.isNotNull(firstRow);
assert.isEmpty(firstRow.querySelectorAll('td'));
```

يجب أن يكون لديك صفان على الأقل يحملان الصنف `sharp`.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="sharp"]').length, 2);
```

يجب أن يكون لديك صفان على الأقل يحملان الصنف `half`.

```js
assert.isAtLeast(document.querySelectorAll('tr[class~="half"]').length, 2);
```

عناصر `th` في العمود الأول يجب أن تحمل `class` باسم `time` وتحتوي على قيم الوقت.

```js
const firstColumnData = [...document.querySelectorAll('tr th:first-child')];
assert.isNotEmpty(firstColumnData);
for (let cell of firstColumnData) {
    if (cell.innerText.match(/\d/)?.length) {
        assert.isTrue(cell.classList.contains('time'));
    }
}
```

يجب أن تحتوي على أربعة عناصر `th` على الأقل من صنف `time` تحتوي على قيم زمنية.

```js
const timeClass = document.querySelectorAll('.time');
assert.isAtLeast(timeClass.length, 4);
for (let cell of timeClass) {
    assert.match(cell.innerText, /\d/)
}
```

يجب أن تحمل جميع عناصر `td` لديك `class` باسم `available-#`، حيث `#` رقم من `0` إلى `5`.

```js
const cells = [...document.querySelectorAll('tr[class="sharp"]>td, tr[class="half"]>td')];
assert.isNotEmpty(cells);
for (let cell of cells) {
    const classString = [...cell.classList].join(' ');
    assert.lengthOf(classString.match(/(?<=\s|^)available-[0-5](?=\s|$)/gm), 1);
}
```

يجب أن يكون لديك ثمانية عناصر `.available-#` على الأقل، حيث `#` رقم بين `0` و `5`.

```js
assert.isAtLeast(document.querySelectorAll('td[class*="available-"]').length, 8);
```

يجب تعريف متغير `--color0` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color0'))
```

يجب تعريف متغير `--color1` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color1'))
```

يجب تعريف متغير `--color2` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color2'))
```

يجب تعريف متغير `--color3` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color3'));
```

يجب تعريف متغير `--color4` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color4'));
```

يجب تعريف متغير `--color5` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--color5'));
```

يجب استخدام `--color0` لتعيين لون الخلفية لعناصر `.available-0`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-0');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color0)');
```

يجب استخدام `--color1` لتعيين لون الخلفية لعناصر `.available-1`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-1');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color1)');
```

يجب استخدام `--color2` لتعيين لون الخلفية لعناصر `.available-2`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-2');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color2)');
```

يجب استخدام `--color3` لتعيين لون الخلفية لعناصر `.available-3`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-3');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color3)');
```

يجب استخدام `--color4` لتعيين لون الخلفية لعناصر `.available-4`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-4');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color4)');
```

يجب استخدام `--color5` لتعيين لون الخلفية لعناصر `.available-5`.

```js
const a = new __helpers.CSSHelp(document).getStyle('.available-5');
assert.exists(a); 
assert.strictEqual(a.backgroundColor, 'var(--color5)');
```

يجب تعريف متغير `--solid-border` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--solid-border'));
```

يجب تعريف متغير `--dashed-border` في محدد الجذر.

```js
const root = new __helpers.CSSHelp(document).getStyle(':root');
assert.exists(root); 
assert.isNotEmpty(root.getPropVal('--dashed-border'));
```

يجب استهداف عناصر `td` التي هي أبناء لعناصر `.sharp`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
```

يجب استخدام `--solid-border` لتعيين `bottom-border` لعناصر `td` التي هي أبناء لعناصر `.sharp`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.sharp td') || new __helpers.CSSHelp(document).getStyle('.sharp>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--solid-border)');
```

يجب استهداف عناصر `td` التي هي أبناء لعناصر `.half`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
```

يجب استخدام `--dashed-border` لتعيين `bottom-border` لعناصر `td` التابعة لعناصر `.half`.

```js
const tds = new __helpers.CSSHelp(document).getStyle('.half td') || new __helpers.CSSHelp(document).getStyle('.half>td');
assert.exists(tds);
assert.strictEqual(tds.borderBottom, 'var(--dashed-border)');
```

يجب أن يكون لديك عنصر `div` يحمل `id` باسم `legend`.

```js
assert.equal(document.querySelector('#legend')?.tagName, 'DIV');
```

يجب أن يحتوي `#legend` على عنصر `span` بالنص `Availability`.

```js
assert.equal(document.querySelector('#legend span')?.textContent.trim(), 'Availability');
```

يجب أن يكون لديك عنصر `div` بالـ `id` الخاص بـ `legend-gradient` داخل عنصر `#legend`.

```js
assert.equal(document.querySelector('#legend #legend-gradient')?.tagName, 'DIV');
```

يجب تعيين صورة الخلفية لعنصر `#legend-gradient` إلى تدرج خطي.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient); 
assert.include(legendGradient.backgroundImage, "linear-gradient(")
```

يجب أن تكون الانتقالات من لون إلى اللون التالي بخط حاد في `#legend-gradient`. تأكد من أن النسب المئوية أعداد صحيحة.

```js
const legendGradient = new __helpers.CSSHelp(document).getStyle('#legend-gradient');
assert.exists(legendGradient);
const gradientStr = legendGradient.backgroundImage;

const COLOR_COUNT = 6;
const EXPECTED_SHORTHAND_STOPS = 6;
const EXPECTED_EXPANDED_STOPS = 12;

for (let i = 0; i < COLOR_COUNT; i++) {
  assert.include(gradientStr, `--color${i}`, `gradient should include --color${i}`);
}

const GRADIENT_PATTERNS = {
  SHORTHAND: /var\(\s*--color[0-5]\s*\)\s+\d+%\s+\d+%/g,
  EXPANDED: /var\(\s*--color[0-5]\s*\)\s+\d+%(?!\s+\d+%)/g
};

const matchGradientFormat = (gradientStr, pattern) => gradientStr.match(pattern);

const shorthandMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.SHORTHAND);
const expandedMatches = matchGradientFormat(gradientStr, GRADIENT_PATTERNS.EXPANDED);

const colorPattern = /var\(\s*--color(\d)\s*\)/g;
const allColorMatches = [...gradientStr.matchAll(colorPattern)];
const colorNumbers = allColorMatches.map(match => parseInt(match[1]));

if (shorthandMatches?.length === EXPECTED_SHORTHAND_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 1, `In shorthand format, --color${i} should appear exactly once, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const idx = colorNumbers.indexOf(i);
    assert.isAtLeast(idx, 0, `--color${i} should be in the gradient`);
    if (i > 0) {
      const prevIdx = colorNumbers.indexOf(i - 1);
      assert.isBelow(prevIdx, idx, `--color${i - 1} should appear before --color${i}`);
    }
  }
} else if (expandedMatches?.length === EXPECTED_EXPANDED_STOPS) {
  for (let i = 0; i < COLOR_COUNT; i++) {
    const count = colorNumbers.filter(num => num === i).length;
    assert.equal(count, 2, `In expanded format, --color${i} should appear exactly twice, but it appears ${count} times`);
  }
  
  for (let i = 0; i < COLOR_COUNT; i++) {
    const firstIdx = colorNumbers.indexOf(i);
    const lastIdx = colorNumbers.lastIndexOf(i);
    assert.equal(lastIdx - firstIdx, 1, `--color${i} should appear in consecutive positions for hard lines`);
    
    if (i > 0) {
      const prevLastIdx = colorNumbers.lastIndexOf(i - 1);
      assert.isBelow(prevLastIdx, firstIdx, `--color${i - 1} should complete before --color${i} starts`);
    }
  }
  
  const percentPattern = /var\(\s*--color\d\s*\)\s+(\d+)%/g;
  const percentMatches = [...gradientStr.matchAll(percentPattern)];
  const percentages = percentMatches.map(match => parseInt(match[1]));
  
  for (let i = 1; i < percentages.length - 1; i += 2) {
    assert.equal(percentages[i], percentages[i + 1], `For hard lines, percentage at position ${i} (${percentages[i]}%) should equal percentage at position ${i + 1} (${percentages[i + 1]}%) where colors meet`);
  }
} else {
  assert.fail('Gradient must use either shorthand format (6 stops with two percentages each) or expanded format (12 stops with one percentage each)');
}
```

يجب أن يحتوي جدولك على عنصر `caption`.

```js
const caption = document.querySelector('table caption');
assert.isNotNull(caption);
```

يجب أن يحتوي عنصر `caption` على النص `Weekly Meeting Availability Schedule`.

```js
const caption = document.querySelector('table caption');
assert.equal(caption?.textContent, 'Weekly Meeting Availability Schedule');
```

يجب أن تحتوي جميع خلايا القسم العلوي في الصف الأول على خاصية `scope` بقيمة `col`.

```js
const headers = [...document.querySelectorAll('tr:first-child>th:not([class*="time"])')];
assert.isNotEmpty(headers);
headers.forEach(header => {
    assert.equal(header.getAttribute('scope'), 'col');
});
```

يجب أن تحتوي جميع خلايا القسم العلوي في العمود الأول على خاصية `scope` بقيمة `row`.

```js
const timeHeaders = [...document.querySelectorAll('th.time')];
assert.isNotEmpty(timeHeaders);
timeHeaders.forEach(header => {
    assert.equal(header.getAttribute('scope'), 'row');
});
```

يجب أن تحتوي جميع خلايا البيانات التي تحمل صنف التوفر على خاصية `aria-label`.

```js
const dataCells = [...document.querySelectorAll('td[class*="available"]')];
assert.isNotEmpty(dataCells);
dataCells.forEach(cell => {
    assert.isNotNull(cell.getAttribute('aria-label'));
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-0` على خاصية `aria-label` بقيمة `0 available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-0')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '0 available');
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-1` على خاصية `aria-label` بقيمة `1 available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-1')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '1 available');
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-2` على خاصية `aria-label` بقيمة `2 available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-2')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '2 available');
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-3` على خاصية `aria-label` بقيمة `3 available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-3')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '3 available');
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-4` على خاصية `aria-label` بقيمة `4 available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-4')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '4 available');
});
```

يجب أن تحتوي خلايا البيانات التي تحمل الصنف `available-5` على خاصية `aria-label` بقيمة `5+ available` إذا كانت موجودة.

```js
const cells = [...document.querySelectorAll('td.available-5')];
cells.forEach(cell => {
    assert.equal(cell.getAttribute('aria-label'), '5+ available');
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Availability Table</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <table>
            <caption>Weekly Meeting Availability Schedule</caption>
            <tbody>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">9 AM</th>
                    <th scope="col">Monday</th>
                    <th scope="col">Tuesday</th>
                    <th scope="col">Wednesday</th>
                    <th scope="col">Thursday</th>
                    <th scope="col">Friday</th>
                </tr>
                <tr class="half">
                    <td class="available-0" aria-label="0 available"></td>
                    <td class="available-1" aria-label="1 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">10 AM</th>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>

                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">11 AM</th>
                    <td class="available-0" aria-label="0 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-0" aria-label="0 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-1" aria-label="1 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">12 PM</th>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-0" aria-label="0 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-0" aria-label="0 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">1 PM</th>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-0" aria-label="0 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">2 PM</th>
                    <td class="available-0" aria-label="0 available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">3 PM</th>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">4 PM</th>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                </tr>
                <tr class="half">
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-0" aria-label="0 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                    <td class="available-2" aria-label="2 available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                </tr>
                <tr class="sharp">
                    <th rowspan="2" class="time" scope="row">5 PM</th>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-5" aria-label="5+ available"></td>
                    <td class="available-3" aria-label="3 available"></td>
                    <td class="available-4" aria-label="4 available"></td>
                </tr>
                <tr>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                    <td></td>
                </tr>
            </tbody>
        </table>
        <div id="legend">
            <span id="legend-title">Availability</span>
            <div id="legend-line"><span>0</span>
                <div id="legend-gradient"></div><span>5+</span>
            </div>
        </div>
    </main>
</body>

</html>
```

```css
:root {
    font-size: 16px;
    font-family: Arial, Helvetica, sans-serif;
    --color0: hsl(0, 0%, 100%);
    --color1: hsl(120, 95%, 90%);
    --color2: hsl(120, 95%, 80%);
    --color3: hsl(120, 95%, 65%);
    --color4: hsl(120, 95%, 40%);
    --color5: hsl(120, 95%, 25%);
    --solid-border: 0.1rem solid black;
    --dashed-border: 0.09rem dashed black;
}

* {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

body {
    background-color: hsl(150, 30%, 75%);
    width: 100vw;
    height: 100vh;
}

main {
    width: 100%;
    height: 100%;
    display: flex;
    flex-direction: column;
    align-items: center;
    justify-content: space-evenly;
}

table {
    border-collapse: collapse;
    table-layout: fixed;
    width: 32rem;
}

tr {
    height: 1.1rem;
}

tr :first-child {
    width: 4rem;
}

td {
    border-right: var(--solid-border);
    border-left: var(--solid-border);
}

.sharp td {
    border-bottom: var(--solid-border);
}

.sharp th:not([class*="time"]) {
    border-bottom: var(--solid-border);
}

.half td {
    border-bottom: var(--dashed-border);
}

tbody :last-child td {
    border: 0;
}

.time {
    text-align: right;
    border: 0;
    padding: 0.4rem;
}

.available-0 {
    background-color: var(--color0);
}

.available-1 {
    background-color: var(--color1);
}

.available-2 {
    background-color: var(--color2);
}

.available-3 {
    background-color: var(--color3);
}

.available-4 {
    background-color: var(--color4);
}

.available-5 {
    background-color: var(--color5);
}

#legend {
    display: flex;
    flex-direction: column;
    align-items: center;
    width: 18rem;
    height: 3.5rem;
    text-align: center;
}

#legend span {
    display: inline-block;
    width: 5rem;
    height: 2rem;
}

#legend-gradient {
    width: 100%;
    height: 60%;
    border: var(--solid-border);
    background-image: linear-gradient(90deg,
            var(--color0) 0% 17%,
            var(--color1) 17% 34%,
            var(--color2) 34% 50%,
            var(--color3) 50% 67%,
            var(--color4) 67% 83%,
            var(--color5) 83% 100%);
}

#legend-line {
    width: 60%;
    height: 100%;
    display: flex;
}
```
