---
id: 69a00a051b9a18694d00c16b
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、連結リストを実装します。

まず、`initList` という名前の関数を作成してください。その中で、`head` プロパティが `null` に設定されているオブジェクトと、`length` プロパティが `0` のオブジェクトを返してください。

これは連結リストを初期化するために呼び出す関数です。`head` はリストの最初の要素を表し、最初はリストに要素がないため `null` です。`length` はリスト内のノード（要素）の数を追跡し、ノードを追加または削除するときに増減させます。

# --hints--

`initList` 関数を作成してください。

```js
assert.isFunction(initList)
```

`initList` 関数はオブジェクトを返す必要があります。

```js
assert.isObject(initList());
```

`initList` 関数は `head` が `null` に設定されたオブジェクトを返す必要があります。

```js
assert.propertyVal(initList(), "head", null);
```

`initList` 関数は `length` が `0` に設定されたオブジェクトを返す必要があります。

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
