---
id: 69a03f70b50b4c4d9a0f0113
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

ループの中で、`previous` を `current` に、`current` を `current.next` に代入し直してください。これにより、`current` は次のノードをチェックし、`previous` は常に1つ前のノードにとどまります。

# --hints--

`previous` 変数を `current` に代入し直す必要があります。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /previous\s*=\s*current;?/);
```

`current` 変数を `current.next` に代入し直す必要があります。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) { 
        --fcc-editable-region--
                
        --fcc-editable-region--
    }
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
