---
id: 69a00eaa43458e618331873c
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

これからリストにノードを追加する関数に取り組みます。

`add` 関数を `list` と `element` のパラメータで作成してください。

# --hints--

`add` 関数を作成してください。

```js
assert.isFunction(add)
```

`add` 関数は `list` と `element` のパラメータを持つべきです。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
