---
id: 69a014bd9b1730e94136279f
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

リストが空でない場合は、その末尾を見つけて新しい要素を追加し、`null`を指すようにしたいです。つまり、`head`からリストをループする必要があります。

それを始めるには、`else`節を開いてください。その中で、`let`を使って`current`変数を作成し、`list.head`に設定します。

# --hints--

`else`節があるはずです。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

`current`節の中で、`list.head`に設定された`else`変数を作成する必要があります。

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
