---
id: 69a04466917c66b132323a1b
title: ステップ 18
challengeType: 1
dashedName: step-18
---

# --description--

もし `previous` が `null` でなければ、これは削除すべき要素がリストの先頭ではなく途中にあることを意味します。そのため、現在のノードは前のノードを次のノードに直接つなげてバイパスする必要があります。

そのために、別の `if` 文を作成し、previous が厳密に `null` でないかどうかをチェックしてください。もしそうであれば、`previous.next` を `current.next` に設定します。

# --hints--

条件 `if` を持つ別の `previous !== null` 文を作成してください。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

`previous.next` を `current.next` に設定してください。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
