---
id: 69a03aaa1e6fe8152788ce6a
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

`remove` 関数の中で、`let` を使って `previous` に設定された `null` 変数を作成し、さらに `let` を使って `current` に設定された `list.head` 変数を作成してください。これらの変数はリストをたどるループで使用します。

`current` はリスト内で現在チェックしているノードを追跡し、先頭から始まり一つずつ前に進みます。一方、`previous` は現在のノードの直前のノードを追跡します。

ノードを削除するときに重要なのは、前のノードの next 値を更新して、削除するノードの次のノードを指すようにし、リストのつながりを保つことです。

# --hints--

`let` 関数の中で `previous` を使って `null` に設定された `remove` 変数を作成してください。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

`let` 関数の中で `current` を使って `list.head` に設定された `remove` 変数を作成してください。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
