---
id: 69a030fbe487f185f55b8627
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

これから、リストからノードを削除する関数に取り組みます。

`remove` 関数を `list` と `element` のパラメータで作成してください。

# --hints--

`remove` 関数を作成してください。

```js
assert.isFunction(remove)
```

`remove` 関数は `list` と `element` のパラメータを持つべきです。

```js
const removeParams = __helpers.getFunctionParams(remove.toString());
const removeParamsNames = removeParams.map(param => param.name)
assert.include(removeParamsNames, "list");
assert.include(removeParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
