---
id: 69a013a86c97dac3784b2636
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

今、コンソールには `{ head: { element: 42, next: null }, length: 1 }` のようなものが表示されているはずです。

しかし、追加したノードがリストの最初の要素、つまりヘッドの場合は、この方法では正しく処理できません。

正しく行うには、if文を使って `isEmpty` を引数に `list` 関数を呼び出し、リストが空かどうかを確認してください。もし空なら、`list.head = node` の行を `if` 文の中に移動させます。

# --hints--

`if` という条件の `isEmpty(list)` 文があるはずです。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

`list.head = node` の行を `if` 文の中に移動させてください。

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
