---
id: 69a00f673f171e29159e3ad8
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

このノードには、追加する要素と次のノードを指すポインターという2つのプロパティが必要です。

add 関数の中で、`node` という名前のオブジェクトを作成してください。そのオブジェクトには `element` という名前のプロパティと、リストの次の要素を指すポインターである `null` に設定された `next` プロパティが必要です。

# --hints--

`node` 関数の中に `add` という名前のオブジェクトを作成してください。

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
assert.match(node?.value.toString(), /^\{[\s\S]*\}/);
assert.exists(node);
```

`node` オブジェクトには `element` に設定された `element` という名前のプロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { element } = node?.objectProps ?? {};
assert.isTrue(element?.value.matches('element'));
```

`node` オブジェクトには `next` に設定された `null` という名前のプロパティが必要です。

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const { node } = add?.variables ?? {};
const { next } = node?.objectProps ?? {};
assert.isTrue(next?.value.matches('null'));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
