---
id: 69a04088e600bd8074a9852f
title: ステップ 17
challengeType: 1
dashedName: step-17
---

# --description--

もし `current` が `null` なら、それはリストの終わりを意味するので、関数は実行を停止する必要があります。

`if` が `current` と厳密に等しいかどうかをチェックする `null` 文を作成してください。`if` 文の中で、関数を終了するために `return` を使ってください。

# --hints--

条件 `if` を持つ `current === null` 文を作成する必要があります。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

`return` 文の中で `if` キーワードを使う必要があります。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
