---
id: 69a0179a85b00b6127e0af45
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`while` ループの中で、`current` を `current.next` に更新して、リストを1つずつノードを進み、`next` プロパティが `null` のノード、つまりリストの終わりに到達したら停止するようにしてください。

# --hints--

`current` ループの中で `current.next` 変数の値を `while` に更新する必要があります。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
