---
id: 69a01864d343dc2f2484d131
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

最後のノードを見つけたら、新しいノードをそこに接続する必要があります。そのために、`current.next` 節の中で、whileループの下にある `node` を `else` に設定してください。

# --hints--

`current.next` ループの下に `node` を `while` に設定する必要があります。

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
