---
id: 69a012c6224fac85832247dd
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

コンソールに表示されている `true`、`{ head: null, length: 0 }`、`true` の結果に気づきましたか？これは、リスト `myList` に何も代入しておらず、`length` プロパティをインクリメントしていないために起こっています。

さあ、`node` を `list.head` に代入し、その下で `length` の `list` プロパティをインクリメントしてください。

# --hints--

`node` を `list.head` に代入する必要があります。

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.head=node');
```

`length` の `list` プロパティをインクリメントしてください。

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'list.length++');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
