---
id: 69a0102e761a245bb1cbd8e7
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

この時点で、連結リストをテストしてより理解を深めましょう。

`myList` という変数を作成し、`initList()` に設定してください。その下に、`isEmpty()` をパラメータとして使い、`myList` の呼び出しをコンソールにログしてください。その後、`add` と任意の数字を使って `myList` を呼び出し、`myList` をコンソールにログしてください。その下に、再び `isEmpty()` をパラメータとして使い、`myList` の呼び出しをコンソールにログしてください。

# --hints--

`myList` という変数を作成し、`initList()` に設定してください。

```js
assert.deepEqual(myList, initList());
```

`isEmpty(myList)` の呼び出しをコンソールにログしてください。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(isEmpty(myList))');
```

`add` を `myList` と任意の数字をパラメータとして呼び出してください。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedCode, /add\(myList,\d+\)/);
```

`myList` 変数をコンソールにログしてください。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(myList)');
```

再び `isEmpty(myList)` をコンソールにログしてください。

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const regexMatch = strippedCode.matchAll(/console\.log\(isEmpty\(myList\)\)/g);
assert.isAtLeast([...regexMatch].length, 2);
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
