---
id: 69a03ddce82b0626f329b001
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

今、`while` ループを作成して、`current` が厳密に `null` でなく、かつ `current.element` が要素と厳密に等しくないかをチェックしてください。

これは、リストの終わりに到達するか、削除する要素が見つかるまでループがリストをたどり続けることを意味します。

# --hints--

条件 `while` を持つ `current !== null && current.element !== element` ループを作成してください。

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
