---
id: 69a00cd4ec513345d40e95ec
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

次に取り組むのは、リストが空かどうかを確認できる関数です。

そのために、`isEmpty` パラメータを受け取る `list` という名前の関数を作成してください。関数の中で、`length` の `list` プロパティが `0` と厳密に等しいかどうかをチェックして返してください。

# --hints--

`isEmpty` 関数を作成してください。

```js
assert.isFunction(isEmpty)
```

`isEmpty` 関数は `list` パラメータを持つべきです。

```js
const regex = __helpers.functionRegex('isEmpty', ['list']);
assert.match(__helpers.removeJSComments(code), regex);
```

`isEmpty` 関数は `list.length === 0` を返すべきです。

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
