---
id: 60a3e3396c7b40068ad6996f
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

`div` 안에 `body` 요소를 추가하세요.

`class` 속성을 `canvas`와 같게 설정하세요.

이것이 그림을 그릴 캔버스 역할을 합니다.

# --hints--

코드에 `div` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('div').length === 1)
```

`div` 요소는 값이 `class`인 `canvas`를 가져야 합니다.

```js
assert(document.querySelector('div').className.split(' ').includes('canvas'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
