---
id: 60a3e3396c7b40068ad6996c
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

이전 CSS 박스 모델 수업에서 모든 HTML 요소가 네 개의 영역을 가진 박스로 취급된다는 것을 배웠습니다.

좋아하는 온라인 소매점에서 박스를 받았다고 상상해 보세요 — 내용물은 박스 안의 아이템이며, 여러분의 경우 헤더, 문단, 이미지 요소입니다.

`src`의 `<img>` 속성을 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`에서 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`으로 변경하세요.

# --hints--

`img` 요소는 `src`인 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`를 가져야 합니다.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```
