---
id: 60a3e3396c7b40068ad69987
title: Step 28
challengeType: 0
dashedName: step-28
---

# --description--

요소 크기를 지정할 때 항상 픽셀을 사용할 필요는 없습니다.

새 규칙 `.three`을 만들고 `width` 속성을 `91%`로 설정하세요.

# --hints--

`.three` 선택자를 사용하세요.

```js
const hasThree = new __helpers.CSSHelp(document).getStyle('.three');
assert(hasThree);
```

`width` 속성을 `91%`로 설정하세요.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '91%');
assert(hasWidth);
```

`.three` 요소는 `width` 값이 `91%`이어야 합니다.

```js
const threeWidth = new __helpers.CSSHelp(document).getStyle('.three')?.getPropertyValue('width');
assert(threeWidth === '91%');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: auto;
}

--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
