---
id: 60a3e3396c7b40068ad69970
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

추가한 `div`에 스타일을 적용하기 전에 CSS를 HTML에 연결해야 합니다.

`link` 요소를 추가해 `styles.css` 파일을 연결하세요. `href`를 `styles.css`로 설정하고, `rel` 속성을 `stylesheet`로 설정하는 것을 잊지 마세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert(/<link/.test(code))
```

`link` 요소는 `rel` 속성을 가지고, 그 값은 `stylesheet`이어야 합니다.

```js
assert.match(code, /<link.*?rel=('|"|`)stylesheet\1/)
```

`link` 요소는 `href` 속성을 가지고, 그 값은 `styles.css`이어야 합니다.

```js
assert.match(code, /<link.*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
