---
id: 60a3e3396c7b40068ad6998c
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

그림의 색상과 모양이 너무 선명해서 Rothko 작품처럼 보이지 않습니다.

`filter` 속성을 사용해 `blur` 요소에서 `2px`하여 그림을 `.canvas`하세요.

다음은 3px `blur`를 추가하는 규칙 예시입니다:

```css
p {
  filter: blur(3px);
}
```

# --hints--

`filter` 속성을 `blur(2px)`로 설정하세요.

```js
const hasFilter = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.filter === 'blur(2px)');
assert(hasFilter);
```

`.canvas` 요소는 `filter` 값을 가진 `blur(2px)`이어야 합니다.

```js
const canvasFilter = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('filter');
assert(canvasFilter === 'blur(2px)');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
--fcc-editable-region--
  
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}

.two {
  width: 475px;
  height: 200px;
  background-color: #8f0401;
  margin: 0 auto 20px;
}

.three {
  width: 91%;
  height: 28%;
  background-color: #b20403;
  margin: auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
        <div class="two"></div>
        <div class="three"></div>
      </div>
    </div>
  </body>
</html>
```
