---
id: 60a3e3396c7b40068ad69978
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

여백을 사용해 요소 바깥쪽 간격을 조절하세요.

`margin` 속성을 사용해 `.frame` 요소에 수직 여백을 `20px`, 수평 여백을 `auto`로 지정하세요. 이렇게 하면 프레임이 20픽셀 아래로 내려가고 페이지에서 수평으로 가운데에 위치합니다.

# --hints--

`margin` 속성을 `20px auto`로 설정해야 합니다.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.margin === '20px auto');
assert(hasMargin);
```

`.frame` 요소는 `margin` 값이 `20px auto`이어야 합니다.

```js
const frameMargin = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('margin');
assert(frameMargin === '20px auto');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
