---
id: 60a3e3396c7b40068ad69979
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

`div` 요소 안에 새로운 `.canvas` 요소를 추가하세요.

새로운 `div`에 `class` 속성을 주고 값은 `one`로 하세요. 이것이 첫 번째 사각형이 됩니다.

# --hints--

새로운 `div` 요소를 만들어야 합니다.

```js
assert(document.querySelectorAll('div').length === 3);
```

새로운 `div` 요소를 `.canvas` 요소 안에 중첩하세요.

```js
assert(document.querySelector('.canvas').children[0].tagName === 'DIV');
```

새로운 `div`에는 `class` 속성이 있고 값은 `one`이어야 합니다.

```js
assert(document.querySelector('.canvas').children[0].className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
