---
id: 60a3e3396c7b40068ad6997f
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

캔버스의 위, 아래, 왼쪽, 오른쪽에 1픽셀 패딩을 추가하니 크기가 502픽셀 x 602픽셀로 변경되었습니다.

캔버스를 원래 크기로 되돌리려면 `padding` 속성을 제거하고 `overflow` 속성을 `hidden`로 설정하세요.

# --hints--

`padding` 선택자에서 `.canvas` 속성을 제거해야 합니다.

```js
const canvasPadding = new __helpers.CSSHelp(document).getStyle('.canvas').getPropertyValue('padding');
assert(!canvasPadding);
```

`overflow` 속성을 `hidden`로 설정해야 합니다.

```js
const hasOverflow = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.overflow === 'hidden');
assert(hasOverflow);
```

`.canvas` 요소는 `overflow` 값이 `hidden`여야 합니다.

```js
const canvasOverflow = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('overflow');
assert(canvasOverflow === 'hidden')
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
--fcc-editable-region--
  padding: 1px;
--fcc-editable-region--
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
      </div>
    </div>
  </body>
</html>
```
