---
id: 60a3e3396c7b40068ad6996b
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

프로젝트를 시작하려면 `img` 요소 안에 `src` 값을 가진 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 속성을 가진 `body` 요소를 추가하세요.

# --hints--

코드에 `img` 요소가 있어야 합니다.

```js
assert(document.querySelectorAll('img').length === 1);
```

`img` 요소는 `src` 값을 가진 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` 속성을 가져야 합니다.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
