---
id: 60a3e3396c7b40068ad69980
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

`one` 요소 바로 아래에 `class` 값이 `two`인 또 다른 `div`를 추가하세요. 이것이 두 번째 사각형이 됩니다.

# --hints--

기존 `.one` 요소는 변경하지 마세요.

```js
assert(document.querySelectorAll('.one').length === 1);
```

`.canvas` 요소 안에 두 번째 `div` 요소가 있어야 합니다.

```js
assert(document.querySelector('.canvas').children[1].tagName === 'DIV');
```

두 번째 `div` 요소는 `class` 값이 `two`여야 합니다.

```js
assert(document.querySelector('.canvas').children[1].className.split(' ').includes('two'));
```

`.two` 요소는 `.one` 요소 다음에 와야 합니다.

```js
assert(document.querySelector('.two').previousElementSibling.className.split(' ').includes('one'));
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
  overflow: hidden;
}

.frame {
  border: 50px solid black;
  width: 500px;
  padding: 50px;
  margin: 20px auto;
}

.one {
  width: 425px;
  height: 150px;
  background-color: #efb762;
  margin: 20px auto;
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
        <div class="one"></div>
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
