---
id: 60a3e3396c7b40068ad6996d
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

내용은 패딩이라고 하는 공간으로 둘러싸여 있는데, 이는 뽁뽁이가 물건과 그 주위 상자 사이를 분리하는 것과 비슷합니다.

테두리는 물건이 배송된 골판지 상자처럼 생각하세요.

`src` 속성을 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`으로 변경하세요

# --hints--

`img` 요소는 `src`인 `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png`을 가져야 합니다

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png">
--fcc-editable-region--
  </body>
</html>
```
