---
id: 60a3e3396c7b40068ad69977
title: Step 13
challengeType: 0
dashedName: step-13
---

# --description--

패딩을 사용해 요소 내부의 간격을 조절하세요.

`.frame`에서 `padding` 단축 속성을 사용해 `.frame`와 `.canvas` 요소 사이의 간격을 `50px`만큼 늘리세요. 이 단축 속성은 요소의 테두리와 캔버스 내부의 위, 아래, 왼쪽, 오른쪽 공간을 모두 늘립니다.

# --hints--

`padding` 속성을 `50px`로 설정해야 합니다.

```js
const hasPadding = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.padding === '50px');
assert(hasPadding);
```

`.frame` 요소는 `padding` 값이 `50px`여야 합니다.

```js
const framePadding = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('padding');
assert(framePadding === '50px');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}

.frame {
  border: 50px solid black;
  width: 500px;
--fcc-editable-region--
  
--fcc-editable-region--
}
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
