---
id: 60a3e3396c7b40068ad69971
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

이제 CSS 차례입니다.

`<div>`에 텍스트가 없더라도, 내용이 있는 박스로 취급됩니다.
`.canvas` 클래스 선택자를 사용하여 `width`를 500픽셀로 설정하는 CSS 규칙을 작성하세요.
다음은 클래스 `card`의 너비를 300픽셀로 설정하는 CSS 규칙입니다:

```css
.card {
  width: 300px;
}
```

# --hints--

코드에 `.canvas` 선택자가 있어야 합니다.

```js
const hasCanvas = new __helpers.CSSHelp(document).getStyle('.canvas');
assert(hasCanvas)
```

`width` 속성을 `500px`으로 설정해야 합니다.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '500px')
assert(hasWidth);
```

`.canvas` 요소의 `width`가 `500px`여야 합니다.

```js
const width = new __helpers.CSSHelp(document).getStyle('.canvas')?.getPropertyValue('width');
assert(width === '500px');
```

# --seed--

## --seed-contents--

```css
--fcc-editable-region--

--fcc-editable-region--

```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="canvas">
    </div>
  </body>
</html>
```
