---
id: 662a6bc12cde72c32fb526f0
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

<dfn>Kiolesura</dfn> ni kama ramani ya darasa la vitu. Kiolesura kina seti ya njia na vigezo ambavyo darasa la vitu linapaswa kutekeleza.

Anza mradi huu kwa kutangaza darasa la vitu tupu linaloitwa `Equation`. Utatumia darasa hili la vitu kufafanua kiolesura, ramani ya usemi wa jumla wa hesabu.

# --hints--

Unapaswa kufafanua darasa jipya la vitu linaloitwa `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
