---
id: 662bdd364bf2cde1487922a9
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Tofauti na lugha nyingine za programu, Python haitekelezi interfaces katika lugha yake ya msingi, lakini maktaba ya kawaida ya Python inakuwezesha kufafanua interfaces kwa njia rahisi.

Kwa mradi huu, utatumia zana kutoka moduli ya `abc`. Kwa hivyo, import moduli hii katika msimbo wako.

# --hints--

Unapaswa ku-import moduli ya `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("import abc")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
    

eq = Equation()
lin_eq = LinearEquation()
```
