---
id: 662bde88dc84f1e249801b1a
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

`ABC` inasimama kwa *Darasa la vitu la msingi la muhtasari*. Darasa la `ABC` linakuwezesha kubadilisha darasa la kawaida kuwa darasa la muhtasari, ambalo ni darasa linalotumika kama mchoro wa darasa halisi.

Badilisha tamko lako la `import` ili kuimport darasa la `ABC` tu kutoka moduli ya `abc`. Unaweza kuimport kitu maalum `x` kutoka moduli `y` kwa kufuata muundo wa import `from y import x`.

Kisha, badilisha darasa lako la `Equation` kuwa darasa la muhtasari kwa kulifanya lirithi kutoka `ABC`.

# --hints--

Unapaswa kuimport `ABC` kutoka moduli ya `abc`.

```js
({ test: () => assert(runPython(`_Node(_code).has_import("from abc import ABC")`)) })
```

Darasa lako la `Equation` linapaswa kurithi kutoka `ABC`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").inherits_from("ABC")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import abc


class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    pass
    
eq = Equation()
lin_eq = LinearEquation()
--fcc-editable-region--
```
