---
id: 662bd552e1c1d2db1b88ba47
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa, fafanua darasa la vitu lingine linaloitwa `LinearEquation` na lifanye kurithi kutoka `Equation`. Utatumia darasa hili la vitu kuwakilisha mlinganyo wa mstari.

# --hints--

Unapaswa kufafanua darasa la vitu linaloitwa `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("LinearEquation")`)) })
```

Darasa lako la vitu la `LinearEquation` linapaswa kurithi kutoka darasa la vitu la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").inherits_from("Equation")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
--fcc-editable-region--

--fcc-editable-region--
```
