---
id: 662bd8260da84bdd5feae419
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Unataka darasa la `LinearEquation` lifanye utekelezaji na siyo kurithi tu njia zote zilizotangazwa ndani ya darasa la `Equation`, ambalo linapaswa kutumika kama kiolesura.

Kwa sasa, darasa la `Equation` ni darasa mzazi tu la `LinearEquation`. Katika hatua zinazofuata utajifunza jinsi ya kulibadilisha kuwa kiolesura rasmi.

Kwa sasa, tengeneza mfano wa `Equation` na uweke thamani yake katika kigezo `eq`, na mfano wa `LinearEquation` na uweke thamani yake katika kigezo `lin_eq`.

# --hints--

Unapaswa kutangaza kigezo `eq` na kuweka thamani yake mfano wa `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("eq = Equation()")`)) })
```

Unapaswa kutangaza kigezo `lin_eq` na kuweka thamani yake mfano wa `LinearEquation`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation()")`)) })
```

# --seed--

## --seed-contents--

```py
class Equation:
    def solve(self):
        pass
        
    def analyze(self):
        pass
        

class LinearEquation(Equation):
    pass
--fcc-editable-region--

--fcc-editable-region--
```
