---
id: 662bd456896f16d9bd03f1a6
title: Hatua ya 2
challengeType: 20
dashedName: step-2
---

# --description--

Ndani ya darasa la `Equation`, fafanua njia mbili mpya za mfano zenye majina `solve` na `analyze`.

# --hints--

Unapaswa kufafanua njia yenye jina `solve` ndani ya darasa la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("solve")`)) })
```

Njia yako ya `solve` inapaswa kuchukua kigezo kimoja, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("solve").has_args("self")`)) })
```

Unapaswa kufafanua njia yenye jina `analyze` ndani ya darasa la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("analyze")`)) })
```

Njia yako ya `analyze` inapaswa kuchukua kigezo kimoja, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("analyze").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Equation:
    pass
--fcc-editable-region--
```
