---
id: 662fbcef5f05e1b84f541a0c
title: Hatua 13
challengeType: 20
dashedName: step-13
---

# --description--

Kila kitu cha mlinganyo kitaanzishwa kwa kupitisha hoja nyingi kama vile viwango vya mlinganyo, kuanzia kiwango cha n cha \\( x \\) hadi kiwango cha sifuri, ikiwa ni pamoja na kiwango kinachowezekana chenye thamani ya `0`.

Kwa mfano, `LinearEquation(4, 5)` itawakilisha mlinganyo \\( 4x + 5 = 0 \\), ambapo `4` ni kiwango cha kiwango cha kwanza (kikubwa hapa) na `5` ni kiwango cha kiwango cha sifuri.

Unahitaji kuangalia kwamba idadi sahihi ya hoja imetumwa kuanzisha kitu cha mlinganyo.

Ndani ya njia ya `__init__`, tengeneza tamko la `if` kuangalia kama urefu wa `args` ni tofauti na idadi ya viwango ambavyo mlinganyo unapaswa kuwa navyo (`degree + 1`). Ikiwa ni hivyo, ondoa `TypeError` na tumia mfuatano wa herufi ufuatao kutoa ujumbe maalum: `f"'{self.__class__.__name__}' object takes {self.degree + 1} positional arguments but {len(args)} were given"`.

Kisha, rekebisha kosa kwa kupitisha `2` na `3` kuanzisha `lin_eq`.

# --hints--

Unapaswa kutengeneza tamko la `if` linaloangalia kama idadi ya viwango vinavyotumika kuanzisha mlinganyo ni tofauti na `degree + 1`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_conditions()[0]
cond.is_equivalent("(self.degree + 1) != len(args)") or cond.is_equivalent("len(args) != (self.degree + 1)")
`)) })
```

Unapaswa kutoa `TypeError` ndani ya tamko jipya la `if` na tumia mfuatano uliotolewa kutoa ujumbe maalum wa kosa.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[0].find_bodies()[0].has_stmt('raise TypeError(f"\\'{self.__class__.__name__}\\' object takes {self.degree + 1} positional arguments but {len(args)} were given")')
`)) })
```

Unapaswa kupitisha `2` na `3` kuanzisha `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("lin_eq = LinearEquation(2, 3)")`)) })
```


# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        pass        
    
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation()
--fcc-editable-region--
```
