---
id: 662fa2e2cf27c09f21f4f5d0
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Katika Python, aina za data hutambuliwa wakati wa utekelezaji wa msimbo (runtime). Kwa hiyo, huna haja ya kubainisha aina ya data ya kigezo unapotangaza. Hata hivyo, unaweza kuweka maelezo ya aina ya data ya kigezo ili kufafanua kuwa kitahifadhi aina fulani ya data kwa kutumia `variable: <data type> = value` au tu `variable: <data type>`. Kumbuka kuwa mfasiri wa Python haufanyi uthibiti wa aina zilizotumika kuweka maelezo ya kigezo, na kawaida unahitaji zana za nje kufanya hivyo.

Ndani ya darasa la `Equation`, tangaza sifa ya kipengele cha darasa `degree`. Usimpe thamani. Badala yake tumia maelezo ya aina ya `int` kuonyesha kuwa itahifadhi nambari kamili ndani ya madarasa halisi.

Baadaye, utatumia sifa hii ya kipengele cha darasa kama sehemu ya mchakato wa uthibitisho wa hoja zinazotumwa kuanzisha vitu vya mlinganyo.

# --hints--

Unapaswa kutangaza sifa ya kipengele cha darasa yenye jina `degree` na kuiweka maelezo ya aina `int` ndani ya darasa la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_variable("degree").is_equivalent("degree: int")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
--fcc-editable-region--
class Equation(ABC):
    def __init__(self):
        pass
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass
--fcc-editable-region--
lin_eq = LinearEquation()
```
