---
id: 662f6d7c92381a3049e4c987
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Kiolesura hakihitaji kufafanua njia za abstrakti tu, bali pia kinaweza kutekeleza njia zitakazorithiwa na madarasa halisi.

Kabla ya kutekeleza `solve` na `analyze` ndani ya darasa la `Equation`, fafanua njia ya `__init__`. Usitumie dekoreta yoyote juu yake.

# --hints--

Unapaswa kufafanua njia ya `__init__` katika darasa lako la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_function("__init__")`)) })
```

Njia yako ya `__init__` inapaswa kuchukua kigezo kimoja, `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
from abc import ABC, abstractmethod


class Equation(ABC):
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
--fcc-editable-region--
```
