---
id: 6639f947d3a1818c9322c64a
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Hatua ya mwisho ya kuthibitisha viwango ni kuangalia kwamba kiwango cha juu kabisa hakilingani na sifuri. Kumbuka kwamba kiwango cha juu kabisa kinapaswa kupitishwa kama hoja ya kwanza wakati wa kuanzisha kitu.

Ongeza tamko la `if` kwa ajili ya hilo na toa `ValueError` ukitumia mfuatano wa herufi ufuatao kutoa ujumbe maalum: `'Highest degree coefficient must be different from zero'`.

# --hints--

Unapaswa kuunda tamko la `if` linalokagua kama kiwango cha kwanza kilichopitishwa kuanzisha mlinganyo ni sawa na sifuri.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_conditions()[0]
cond.is_equivalent("args[0] == 0") or cond.is_equivalent("0 == args[0]") or cond.is_equivalent("not args[0]")
`)) })
```

Unapaswa kutoa `ValueError` ndani ya tamko jipya la `if` na utumie mfuatano wa herufi uliotolewa kurudisha ujumbe maalum wa kosa.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").find_ifs()[2].find_bodies()[0].has_stmt("raise ValueError('Highest degree coefficient must be different from zero')")
`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        pass
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
```
