---
id: 663a32735b317af9812eb0d7
title: Hatua 27
challengeType: 20
dashedName: step-27
---

# --description--

Katika mlinganyo wa mstari wa aina \\( ax + b = 0 \\), mteremko ni tu msingi wa nambari \\( a \\), na kukutana na mhimili wa y ni msingi wa nambari \\( b \\).

<img alt="a plot of a linear function" src="https://cdn.freecodecamp.org/curriculum/python/linear-equation.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Utaitumia njia ya `analyze` kutoa taarifa zaidi kuhusu mlinganyo huo. Ndani ya njia ya `analyze`, fungua misingi ya nambari na uihifadhi katika vigezo `slope` na `intercept`.

Kisha, rudisha kamusi yenye funguo `'slope'` na `'intercept'` na thamani za mteremko na kukutana na mhimili wa y, mtawalia. Baadaye, fanya wito wa `analyze` kwa `lin_eq` na chapisha matokeo.


# --hints--

Unapaswa kufungua thamani zilizohifadhiwa katika sifa ya kipengele ya `coefficients` na kuzihifadhi katika vigezo `slope` na `intercept` ndani ya njia ya `analyze`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("analyze").has_stmt("slope, intercept = self.coefficients.values()")`)) })
```

Njia ya `analyze` inapaswa kurudisha kamusi yenye funguo `'slope'` na `'intercept'` na thamani za mteremko na kukutana na mhimili wa y, mtawalia.

```js
({ test: () => runPython(`
eq = LinearEquation(2.2, 1.5)
a = eq.analyze()
assert a['slope'] == 2.2, "Expected different slope"
assert a['intercept'] == 1.5, "Expected different intercept"
`) })
```

Unapaswa kuita njia ya `analyze` ya kitu chako cha `lin_eq`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(lin_eq.analyze())")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return x
--fcc-editable-region--
    def analyze(self):
        pass

lin_eq = LinearEquation(2, 3)
print(lin_eq)
print(lin_eq.solve())
--fcc-editable-region--
```
