---
id: 663a22ba7420c4d2f7fd2aec
title: Hatua 25
challengeType: 20
dashedName: step-25
---

# --description--

Ni wakati wa kutekeleza njia ya `solve`. Ikiwa na mlinganyo wa mstari katika fomu \\( ax + b = 0 \\), suluhisho ni \\(x = -\frac{b}{a}\\).

Fungua thamani za coefficients zilizohifadhiwa katika sifa ya kipengele ya `coefficients` na uzihifadhi katika vigezo `a` na `b`. Kumbuka kwamba utahitaji kutumia njia ya `.values()`.

Kisha, tangaza kigezo `x`, weka thamani yake kuwa suluhisho la mlinganyo na lirudishe kutoka kwa njia ya `solve`.

# --hints--

Unapaswa kufungua thamani zilizohifadhiwa ndani ya sifa ya kipengele ya `coefficients` na kuzihifadhi katika vigezo `a` na `b`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("a, b = self.coefficients.values()")`)) })
```

Unapaswa kutangaza kigezo kinachoitwa `x` na kuwekea thamani suluhisho la mlinganyo wa mstari.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_stmt("x = -b/a")`)) })
```

Unapaswa kurudisha `x` kutoka kwa njia yako ya `solve`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").find_function("solve").has_return("x")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')                
        equation_string = ' '.join(terms) + ' = 0'
        return equation_string.strip('+')        
    
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
--fcc-editable-region--
    def solve(self):
        pass
--fcc-editable-region--
    
    def analyze(self):
        pass
    
lin_eq = LinearEquation(2, 3)
print(lin_eq)

```
