---
id: 66562f71937f877c66123bbe
title: Hatua 49
challengeType: 20
dashedName: step-49
---

# --description--

Kitu cha kwanza kuonyeshwa juu kabisa ya matokeo kitakuwa aina ya mlinganyo. Ongeza sifa ya darasa la vitu iitwayo `type` kwa darasa la vitu la `Equation` na uifafanue kwa `str`.

Kisha, ongeza tamko lingine la `if` kwenye njia ya `__init_subclass__` ili kuangalia kama madarasa yanayorithi kutoka `Equation` yana sifa ya `type`. Tumia muundo ule ule wa tamko la `if` lililopo kwa mabadiliko yanayofaa.

Mwishowe, ongeza sifa mpya ya darasa la vitu kwenye darasa la `LinearEquation` na kwenye darasa la `QuadraticEquation`. Weka thamani yake kuwa mfuatano wa herufi `'Linear Equation'` na mfuatano wa herufi `'Quadratic Equation'`, mtawaliwa.

# --hints--

Unapaswa kufafanua mabadiliko ya darasa la vitu iitwayo `type` ndani ya darasa la `Equation` na kuifafanua kwa `str`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").has_stmt("type: str")`)) })
```

Unapaswa kuunda tamko la `if` linaloangalia kama `cls` haina sifa ya `type` ndani ya njia ya `__init_subclass__` na kuinua `AttributeError` ukitumia mfuatano wa herufi uliotolewa.

```js
({ test: () => assert(runPython(`
if_str = """
if not hasattr(cls, 'type'):
  raise AttributeError(
    f\\"Cannot create '{cls.__name__}' class: missing required attribute 'type'\\"
  )
"""
_Node(_code).find_class("Equation").find_function("__init_subclass__").has_stmt(if_str)
`)) })
```

Sifa ya `type` ya darasa la `LinearEquation` inapaswa kuwa na thamani `'Linear Equation'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("LinearEquation").has_stmt("type = 'Linear Equation'")`)) })
```

Sifa ya `type` ya darasa la `QuadraticEquation` inapaswa kuwa na thamani `'Quadratic Equation'`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("QuadraticEquation").has_stmt("type = 'Quadratic Equation'")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
import re

--fcc-editable-region--
class Equation(ABC):
    degree: int
  
    def __init__(self, *args):
        if (self.degree + 1) != len(args):
            raise TypeError(
                f"'Equation' object takes {self.degree + 1} positional arguments but {len(args)} were given"
            )
        if any(not isinstance(arg, (int, float)) for arg in args):
            raise TypeError("Coefficients must be of type 'int' or 'float'")
        if args[0] == 0:
            raise ValueError("Highest degree coefficient must be different from zero")
        self.coefficients = {(len(args) - n - 1): arg for n, arg in enumerate(args)}

    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    def __str__(self):
        terms = []
        for n, coefficient in self.coefficients.items():
            if not coefficient:
                continue
            if n == 0:
                terms.append(f'{coefficient:+}')
            elif n == 1:
                terms.append(f'{coefficient:+}x')
            else:
                terms.append(f"{coefficient:+}x**{n}")
        equation_string = ' '.join(terms) + ' = 0'
        return re.sub(r"(?<!\d)1(?=x)", "", equation_string.strip("+"))        

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass
        
class LinearEquation(Equation):
    degree = 1
    
    def solve(self):
        a, b = self.coefficients.values()
        x = -b / a
        return [x]

    def analyze(self):
        slope, intercept = self.coefficients.values()
        return {'slope': slope, 'intercept': intercept}

class QuadraticEquation(Equation):
    degree = 2
--fcc-editable-region--
    def __init__(self, *args):
        super().__init__(*args)
        a, b, c = self.coefficients.values()
        self.delta = b**2 - 4 * a * c

    def solve(self):
        if self.delta < 0:
            return []
        a, b, _ = self.coefficients.values()
        x1 = (-b + (self.delta) ** 0.5) / (2 * a)
        x2 = (-b - (self.delta) ** 0.5) / (2 * a)
        if self.delta == 0:
            return [x1]

        return [x1, x2]

    def analyze(self):
        a, b, c = self.coefficients.values()
        x = -b / (2 * a)
        y = a * x**2 + b * x + c
        if a > 0:
            concavity = 'upwards'
            min_max = 'min'
        else:
            concavity = 'downwards'
            min_max = 'max'
        return {'x': x, 'y': y, 'min_max': min_max, 'concavity': concavity}

def solver(equation):
    if not isinstance(equation, Equation):
        raise TypeError("Argument must be an Equation object")


lin_eq = LinearEquation(2, 3)
quadr_eq = QuadraticEquation(1, 2, 1)

```
