---
id: 66759e32b88fb5459b1e0234
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Njia ya `__init_subclass__` huitwa kila wakati darasa la vitu linaloibua linapoundwa na inaruhusu kubinafsisha madarasa ya watoto. Njia hii hupokea kigezo kinachoitwa kwa mkataba `cls` (kinawakilisha "darasa la vitu"), ambacho kinakilisha darasa jipya la mtoto.

Tengeneza njia ya `__init_subclass__` katika darasa lako la `Equation` na umpe kigezo cha `cls`.

# --hints--

Unapaswa kutengeneza njia ya `__init_subclass__` yenye kigezo cha `cls` katika darasa lako la `Equation`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init_subclass__").has_args("cls")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
    
    def __init__(self):
        pass
--fcc-editable-region--
    
--fcc-editable-region--
    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
