---
id: 6675aaf418b41157f6ccd692
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Ni wakati wa kurudi kwenye njia ya `__init__`. Kulingana na aina ya mlinganyo, utahitaji kupitisha idadi tofauti ya hoja wakati wa kuanzisha.

Ongeza kigezo cha pili `args` kwenye njia hiyo na tumia kiendeshaji `*` ili kuifanya ikubali idadi tofauti ya hoja.

# --hints--

Njia yako ya `__init__` inapaswa kuchukua vigezo viwili, `self`, na `*args`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Equation").find_function("__init__").has_args("self, *args")`)) })
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod


class Equation(ABC):
    degree: int
--fcc-editable-region--
    def __init__(self):
        pass
--fcc-editable-region--
    def __init_subclass__(cls):
        if not hasattr(cls, "degree"):
            raise AttributeError(
                f"Cannot create '{cls.__name__}' class: missing required attribute 'degree'"
            )

    @abstractmethod
    def solve(self):
        pass
        
    @abstractmethod
    def analyze(self):
        pass


class LinearEquation(Equation):
    degree = 1

    def solve(self):
        pass

    def analyze(self):
        pass


lin_eq = LinearEquation()
```
