---
id: acda2fb1324d9b0fa741e6b5
title: Construir uma Ferramenta de Confirmação do Final
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

Neste laboratório, você implementará uma função que verifica se uma string termina com a string alvo fornecida.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `confirmEnding` que recebe dois parâmetros: a string a ser verificada e a string para comparar.
2. A função deve retornar `true` se a primeira string terminar com a segunda string, e `false` caso contrário.
3. Você não deve usar o método `.endsWith()`; em vez disso, use um dos métodos de substring do JavaScript para alcançar isso.

# --hints--

Você deve criar uma função chamada `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` deve receber 2 parâmetros.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian","n")` deve retornar `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation","on")` deve retornar `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor","n")` deve retornar `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` deve retornar `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name","name")` deve retornar `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame","same")` deve retornar `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` deve retornar `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame","game")` deve retornar `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. Não sabemos por quanto tempo mais podemos suportar o nada", "montanha")` deve retornar `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` deve retornar `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Seu código não deve usar o método embutido `.endsWith()` para resolver o laboratório.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
