---
id: 683ecb010f4c6bafd428802c
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将验证一组医疗数据，以确保其符合一组规则。

声明一个名为 `medical_records` 的变量并赋值为空列表。在接下来的步骤中，你将使用它来保存你的医疗数据。

# --hints--

你应该声明一个名为 `medical_records` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

你应该为你的 `medical_records` 变量赋值一个空列表。

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```

