---
id: 6842ba0086201089ea268c92
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

按照你在前面步骤中使用的相同结构，`medical_records` 列表已为你填充了其他患者的数据。欢迎查看。

接下来你将开始编写用于验证数据集的函数。创建一个名为 `validate` 的函数，带有一个参数 `data`。

你想确保你的数据是列表或元组。因此，在 `validate` 函数中，声明一个名为 `is_sequence` 的变量，并赋值为调用 `isinstance`。将 `data` 作为第一个参数传入，将包含 `list` 和 `tuple` 的元组作为第二个参数传入。

# --hints--

你应该有一个名为 `validate` 的函数。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

你的 `validate` 函数应有一个单一的参数 `data`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

你应该在你的 `validate` 函数中有一个名为 `is_sequence` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

你应该将 `isinstance(data, (list, tuple))` 赋值给你的 `is_sequence` 变量。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
