---
id: 6842ff2b3c45aa4a89785ecb
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

紧接着你的 `if` 语句，声明一个名为 `is_invalid` 的变量并将其设置为 `False`。稍后，你将把它用作标记来运行条件式语句。

# --hints--

你的函数中应该有一个名为 `is_invalid` 的变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_invalid")
`) })
```

你应该将 `False` 赋值给你的 `is_invalid` 变量。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_variable("is_invalid").is_equivalent("is_invalid = False")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False

--fcc-editable-region--
    
--fcc-editable-region--
```
