---
id: 6842aee5bcaf4ab8883f8191
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

正如你在之前的课程中学到的，字典是一种保存密钥-值点对的数据结构，其中密钥是不可变的数据类型（例如字串），并且在字典中是唯一的：

```py
person = {
    'name': 'John',
    'age': 33
}
```

`medical_records` 列表将保存字典，每个字典表现一个病人。为你的 `medical_records` 列表添加一个字典，该字典包含密钥 `patient_id` 和字串值 `P1001`。

# --hints--

你的 `medical_records` 列表应包含一个字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

你的字典中应有一个 `patient_id` 密钥，其值为字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
