---
id: 6842d659a4ae7973dabab3b0
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

通过添加以下三个密钥-值点对来完成字典：

- 一个密钥 `diagnosis`，其值为 `Hypertension`
- 一个键 `medications`，其值为列表 `['Lisinopril']`
- 一个键 `last_visit_id`，其值为 `V2301`。

# --hints--

你的字典中应该有一个 `diagnosis` 密钥，其值为字串 `Hypertension`。

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

你的字典中应该有一个 `medications` 密钥，其值为列表 `['Lisinopril']`。

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

你的字典中应有一个 `last_visit_id` 密钥，其值为字串 `V2301`。

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

你的字典中应有一个 `patient_id` 密钥，其值为字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

你的字典中应该有一个 `age` 密钥，其值为整数 `34`。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

你的字典中应该有一个 `gender` 密钥，其值为字串 `Female`。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
