---
id: 6842d5b537c64266420faf62
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

为你的字典添加一个键 `age`，其值为整数 `34`，以及一个键 `gender`，其值为字串 `Female`。不要忘记键值点对之间的逗号。

# --hints--

你的 `medical_records` 列表应包含一个字典。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

你的字典中应该有一个 `age` 密钥，其值为整数 `34`。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

你的字典中应该有一个 `gender` 密钥，其值为字串 `Female`。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

你的字典中应有一个 `patient_id` 密钥，其值为字串 `P1001`。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
