---
id: 6846b0048f71fc4efd786b42
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

`keys()` 方法返回一个包含字典中所有密钥的视图对象：

```py
person = {
   'name': 'John',
   'age': 33
}

print(person.keys()) # dict_keys(['name, 'age'])
```

在你的 `for` 循环内，第一个 `if` 语句之后，创建一个 `if` 语句，当当前字典的密钥集合与 `key_set` 不同时运行。这样可以确保字典中没有缺失或无效的密钥。

在新的 `if` 语句中，打印 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.`（其中 `<dictionary>` 和 `<index>` 应替换为当前迭代中的字典和索引），并将 `is_invalid` 设置为 `True`。

# --hints--

你应该在你的 `for` 循环内有第二个 `if` 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1]
`) })
```

你的新建 `if` 语句应在 `set(dictionary.keys())` 与 `key_set` 不同时运行。

```js
({ test: () => assert(runPython(`
_condition = _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_conditions()[0]

_condition.is_equivalent("set(dictionary.keys()) != key_set") or _condition.is_equivalent("key_set != set(dictionary.keys())")
`)) })
```

你应该在你的新 `if` 语句中打印 `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.`（其中 `<dictionary>` 和 `<index>` 应该被当前迭代中的字典和索引替换）。

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
_test_data = [
  {'testing': 'stuff'},
  {
    'patient_id': 'P1003',
    'age': 29,
    'gender': 'Female',
    'diagnosis': 'Asthma',
    'medications': ['Albuterol'],
    'last_visit_id': 'V2303',
  },
  {'keep testing': 'stuff'}
]
validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
len(out) >= 2 and _first in out and _second in out
`)) })
```

你应该在你的新建 `if` 语句中将 `is_invalid` 设置为 `True`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

--fcc-editable-region--
        
--fcc-editable-region--

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
