---
id: 6849311be42a0431b8e59109
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

你正在查看传递给你的函数的数据是列表还是元组。你仍然需要确保序列中的每一项都是字典。

在你的 `for` 循环内，如果 `dictionary` 中的项不是 `dict` 的实例，打印 `Invalid format: expected a dictionary at position <index>.`（其中 `<index>` 应替换为当前索引）并将 `is_invalid` 设置为 `True`。

# --hints--

你应该在你的 `for` 循环内有一个 `if` 语句。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0]
`) })
```

你的 `if` 语句应查看 `dictionary` 是否不是 `dict` 的实例。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(dictionary, dict)")
`) })
```

你应该在你的 `if` 语句中打印 `Invalid format: expected a dictionary at position <index>.`（其中 `<index>` 应替换为当前索引）。

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
validate([{}, "", {}, [], 0])
print = built_in_print

_first = 'Invalid format: expected a dictionary at position 1.'
_second = 'Invalid format: expected a dictionary at position 3.'
_third = 'Invalid format: expected a dictionary at position 4.'
len(out) >= 3 and _first in out and _second in out and _third in out
`)) })
```

你应该在你的 `if` 语句内将 `is_invalid` 设置为 `True`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False

    for index, dictionary in enumerate(data):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
