---
id: 684ed67f4c678896445f973f
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

`**` 操作符可用于解包字典中的元素，并将它们作为关键字参数传递到函数调用中：

```py
def sum(a, b, c):
    return a + b + c
nums = {'a': 2, 'b': 4, 'c': 1}

print(sum(**nums)) # 7
```

在上面的示例中，`sum(**nums)` 等同于 `sum(a=2, b=4, c=1)`。

在你的代码底部，打印调用 `find_invalid_records` 函数的结果。对于它的参数，使用 `**` 操作符来打包 `medical_records[0]`。

# --hints--

你应该打印 `find_invalid_records(**medical_records[0])`。

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(find_invalid_records(**medical_records[0]))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
--fcc-editable-region--

--fcc-editable-region--
```
