---
id: 684fe1be09b05ad785832546
title: 步骤 29
challengeType: 20
dashedName: step-29
---

# --description--

现在你的正则表达式匹配字母 `p` 后跟一个或多个数字，最后你需要查看的是字串中没有发现额外的字符。

为此你可以使用 `re` 模块中的另一个函数。`fullmatch` 函数在正则表达式模式匹配整个字串时返回一个匹配对象，否则返回 `None`。

```py
import re

book = "Fahrenheit 451"
print(re.fullmatch('\d+', book)) #None

print(re.fullmatch('Fahrenheit \d+', book))
# <re.Match object; span=(0, 14), match='Fahrenheit 451'>
```

将 `search` 调用替换为 `fullmatch` 调用，保持相同的参数。

# --hints--

你应该将 `re.search('p\d+', patient_id, re.IGNORECASE)` 替换为 `re.fullmatch('p\d+', patient_id, re.IGNORECASE)`。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent("constraints = {'patient_id': isinstance(patient_id, str) and re.fullmatch('p\\d+', patient_id, re.IGNORECASE)}")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str) and re.search('p\d+', patient_id, re.IGNORECASE)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
