---
id: 66e45c8140f9fda5c105ae26
title: Створити сторінку флаєра події
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---


# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У тілі документа має бути елемент `header`.
1. Ваш елемент `header` повинен містити зображення для вашої події та `h1` з назвою події, у такому порядку. Ви можете використати це зображення, якщо хочете: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. У тілі документа має бути елемент `main`.
1. Ваш елемент `main` повинен містити щонайменше два елементи `section`, які демонструють функції події.
1. Кожен ваш елемент `section` повинен містити `h2`.
1. Ваше тіло має мати відступи зверху і знизу розміром `50px`.
1. Ваше тіло має мати поля зверху і знизу розміром `0`, а зліва і справа — такі, що центровані.
1. Ширина вашого тіла має бути встановлена відносно ширини вікна перегляду.
1. Ваше тіло має використовувати функцію `calc` для встановлення властивості `min-height` на 100% висоти вікна перегляду за вирахуванням усіх відступів зверху і знизу.
1. У вашому флаєрі має бути щонайменше один елемент `hr`.
1. Ви маєте встановити значення `width` для всіх елементів `hr` і `section` у відсотках відносно їх батьківського елемента.

**Примітка:** Обов’язково підключіть свій файл стилів у HTML і застосуйте CSS.

# --hints--

У елементі `header` має бути елемент `body`.

```js
assert.exists(document.querySelector('body > header'));
```

У `img` першим дочірнім елементом має бути `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

У `h1` другим дочірнім елементом має бути `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

Має бути лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Ваш `h1` має містити назву вашої події.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

У елементі `main` має бути `body`.

```js
assert.exists(document.querySelector('body > main'));
```

У вашому елементі `section` має бути щонайменше два елементи `main`.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

Кожен елемент `section` має містити `h2`.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

Ваші елементи `h2` не повинні бути порожніми.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

Ваш елемент `body` має мати відступи зверху і знизу розміром `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

Ваш елемент `body` має мати поля зверху і знизу розміром `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

Ваш елемент `body` має мати поля зліва і справа розміром `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

Ваш елемент `body` має встановлювати ширину за допомогою `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

Ваш `body` має використовувати `calc` для встановлення властивості `min-height` у значення `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

У вас має бути щонайменше один елемент `hr`.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

Ширина ваших елементів `hr` має встановлюватися у відсотках.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

Ширина ваших елементів `section` має встановлюватися у відсотках.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
