---
id: 655cd899f8de09431eabb40c
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Merge sort è un algoritmo di ordinamento che usa il principio divide et impera per ordinare collezioni di dati. In pratica, "divide" una collezione in sotto-parti più piccole, e "conquista" le sotto-parti ordinandole indipendentemente, poi unisce le sotto-parti ordinate.

Inizia creando una funzione chiamata `merge_sort` che prende il parametro `array`. Questa funzione gestirà l’attività di ordinare una lista di numeri. Usa la parola chiave `pass` nel corpo della funzione.

# --hints--

Dovresti dichiarare una funzione chiamata `merge_sort` con un solo parametro: `array`. Non dimenticare di usare la parola chiave `pass` nel corpo della funzione.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
