---
id: 656639771fed09b5c0e8fe71
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora che hai diviso la lista `array` in due liste separate, continuerai a dividere ogni lista finché ogni elemento non sarà da solo nella propria lista. Una lista con un solo numero è sempre ordinata.

Per farlo, chiama ricorsivamente `merge_sort` all'interno della tua funzione e passa `left_part` come argomento alla chiamata.

# --hints--

Dovresti chiamare `merge_sort` con l'argomento `left_part` alla fine del corpo della tua funzione.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
