---
id: 656627b47bd2d2a4afbc945d
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Nel passaggio precedente, hai ottenuto il punto medio. Puoi usarlo per dividere `array` in due parti e assegnare ciascuna parte a nuove variabili.

Usa la sintassi slice per estrarre la metà sinistra di `array` e assegnarla a una variabile chiamata `left_part`.

# --hints--

Dovresti avere una variabile chiamata `left_part` nella tua funzione `merge_sort`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Dovresti usare la sintassi slice per estrarre la metà sinistra dell'array e assegnarla alla variabile `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
