---
id: 656710d1e0ec62253426db24
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Il loop `while` che hai creato confronta un elemento da `left_part` con un altro in `right_part`, quindi aggiunge l'elemento più piccolo alla lista principale `array`.

Continuerà questa operazione finché non ci saranno più elementi da confrontare. Ma `left_part` potrebbe ancora avere elementi mentre `right_part` non ne ha, e viceversa.

Crea un altro loop `while` che venga eseguito quando `left_array_index` è minore di `len(left_part)`. Nei passaggi successivi, lo userai per copiare gli elementi rimanenti in `left_part` nella lista `array`.

# --hints--

Il tuo loop `while` dovrebbe avere questa condizione: `left_array_index < len(left_part)`. Non dimenticare di usare la parola chiave `pass` nel corpo del loop.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

Dovresti avere la parola chiave `pass` nel corpo del tuo loop `while`

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--
    
--fcc-editable-region--
```
