---
id: 656702f8b4cbd8cbf0a433c6
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Ora devi creare un loop che continui finché ci sono elementi rimanenti sia in `left_part` sia in `right_part`.

Per farlo, crea un loop `while` con due condizioni: una che verifica se `left_array_index` è minore della lunghezza di `left_part` e un'altra che verifica se `right_array_index` è minore della lunghezza di `right_part`.

# --hints--

Il tuo loop `while` dovrebbe controllare se `left_array_index` è minore della lunghezza di `left_part` e se `right_array_index` è minore della lunghezza di `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Usa la parola chiave `pass` nel corpo del loop.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
