---
id: 656706afd65547d22bee0b68
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

All'interno del loop `while`, sostituisci `pass` con un'istruzione `if` che verifica se l'elemento nel `left_part` è minore dell'elemento nel `right_part`.

Usa la parola chiave `pass` nel corpo dell'istruzione `if`.

# --hints--

Dovresti sostituire la parola chiave `pass` con un'istruzione `if` con questa condizione: `left_part[left_array_index] < right_part[right_array_index]`. Aggiungi la parola chiave `pass` al corpo dell'istruzione `if`.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
