---
id: 6569b743630ee592a65a7e2f
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Prima di testare la funzione `merge_sort()`, devi creare un caso base che fermi l'esecuzione della funzione quando la lunghezza di `array` è minore o uguale a `1`.

Questo caso base fermerà la chiamata ricorsiva. Senza di esso, l'operazione di ordinamento per unione continuerebbe a funzionare anche quando la lista è già ordinata o non contiene elementi.

Subito dopo la dichiarazione della funzione, crea un'istruzione `if` con questa condizione: `len(array) <= 1`. All'interno del corpo dell'istruzione `if` aggiungi un'istruzione `return` per fermare l'esecuzione della funzione.

# --hints--

Dovresti creare un'istruzione `if` per verificare se `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

Dovresti avere un'istruzione `return` nel corpo dell'istruzione `if`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
