---
id: 6569b4e0bd29d17d4b53b16c
title: Passaggio 19
challengeType: 20
dashedName: step-19
---

# --description--

Ora, replicherai la stessa logica del loop `while` per `right_part`.

Crea un loop `while` che venga eseguito quando `right_array_index` è minore di `len(right_part)`. All’interno del loop `while`, assegna `right_part[right_array_index]` a `array[sorted_index]`. Alla fine, incrementa `right_array_index` e `sorted_index` di `1`.

# --hints--

Il tuo loop `while` dovrebbe verificare se `right_array_index` è minore di `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

Dovresti assegnare `right_part[right_array_index]` a `array[sorted_index]` nel corpo del loop.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

Dovresti incrementare `right_array_index` di `1` all’interno del loop `while`. Nota che l’incremento deve avvenire dopo aver assegnato l’elemento della parte destra all’array ordinato.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

Dovresti incrementare il valore corrente di `sorted_index` di `1` all’interno del loop `while`. Nota che l’incremento deve avvenire dopo aver assegnato l’elemento della parte destra all’array ordinato.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
