---
id: 6577254891048c8f2c19e961
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Alla fine del corpo della tua funzione, richiama di nuovo la funzione `merge_sort()`. Questa volta, passa `right_part` come argomento alla chiamata della funzione.

# --hints--

Dovresti chiamare `merge_sort()` con l'argomento `right_part` alla fine del corpo della tua funzione.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
