---
id: 6577562501feabdf0984a184
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Dopo aver assegnato l'elemento in `left_part` all'array ordinato, incrementa `left_array_index` di `1`.

# --hints--

Dovresti aggiungere `1` al valore corrente di `left_array_index` dopo aver inserito l'elemento nell'array ordinato.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body()[1]

    cond.is_equivalent("left_array_index += 1") or cond.is_equivalent("left_array_index = left_array_index + 1") or cond.is_equivalent("left_array_index = 1 + left_array_index")
    `)) 
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]

--fcc-editable-region--
```
