---
id: 657f59751d5a5c9b069ae0f3
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Come hai imparato in una lezione precedente, puoi usare la variabile `__name__` per determinare se uno script Python viene eseguito come programma principale oppure se viene importato come modulo (codice scritto in un altro file Python).

Se il valore di `__name__` è impostato su `'__main__'`, significa che lo script corrente è il programma principale e non un modulo.

In questo progetto, userai lo script corrente come programma principale.

Crea un’istruzione `if` che verifichi se il valore di `__name__` è `'__main__'`.
All’interno dell’istruzione `if`, crea una lista chiamata `numbers` e assegnale questi valori: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

La tua istruzione `if` dovrebbe verificare se `__name__` ha valore `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

Dovresti avere una lista `numbers` con questi valori: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
