---
id: 657f3a661730891aa64f3568
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Usa la sintassi slice per estrarre la metà destra di `array` e assegnarla a una variabile chiamata `right_part`.

# --hints--

Dovresti avere una variabile chiamata `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Dovresti usare la sintassi slice per estrarre la metà destra dell'array e assegnarla alla variabile `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
