---
id: 646c5e727e56e743c9aed4a1
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

CSS-Positionierung ermöglicht es Ihnen, festzulegen, wie ein Element im Browser positioniert werden soll. Es gibt eine `position`-Eigenschaft, die Sie auf `static`, `absolute`, `relative`, `sticky` oder `fixed` setzen können.

Sobald Sie die `position`-Eigenschaft des Elements festgelegt haben, können Sie das Element verschieben, indem Sie einen Pixel- oder Prozentwert für eine oder mehrere der `top`-, `right`-, `left`- oder `bottom`-Eigenschaften angeben.

`static` ist die Standardpositionierung für alle Elemente. Wenn Sie sie einem Element zuweisen, können Sie es nicht mit `top`, `right`, `left` oder `bottom` verschieben.

Geben Sie `.cat-head` eine `position`-Eigenschaft mit dem Wert `static`, und setzen Sie dann die `top`- und `left`-Eigenschaften jeweils auf `100px`.


# --hints--

Ihr `.cat-head`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `static` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'static')
```

Ihr `.cat-head` Selektor sollte eine `top` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px')
```

Ihr `.cat-head` Selektor sollte eine `left` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {


  
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
--fcc-editable-region--
```
