---
id: 646c5d7057c45f432fcdd46c
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Verwenden Sie einen Klassenselektor, um dem `.cat-head`-Element eine Breite von `205px` und eine Höhe von `180px` zu geben. Geben Sie ihm außerdem einen Rahmen von `1px solid #000` und eine `border-radius` von `46%`.

# --hints--

Sie sollten einen `.cat-head`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Ihr `.cat-head`-Selektor sollte eine `width` mit dem Wert `205px` haben.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width, '205px');
```

Ihr `.cat-head`-Selektor sollte eine `height` mit dem Wert `180px` haben.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height,'180px')
```

Ihr `.cat-head`-Selektor sollte eine `border` mit dem Wert `1px solid #000` haben.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border, '1px solid rgb(0, 0, 0)')
```

Ihr `.cat-head`-Selektor sollte eine `border-radius` mit dem Wert `46%` haben.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius, '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
