---
id: 646ceb843412c74edee27a79
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

Verwenden Sie einen Klassenselektor, um dem `.cat-right-ear`-Element die `height`- und `width`-Eigenschaften mit dem Wert `100px` zu geben. Setzen Sie die `background-color` auf `white`. Setzen Sie die linken und rechten Rahmen auf `35px solid blue`. Setzen Sie die oberen und unteren Rahmen auf `35px solid red`.

# --hints--

Sie sollten einen `.cat-right-ear`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ihr `.cat-right-ear`-Selektor sollte eine `height`-Eigenschaft mit dem Wert `100px` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

Ihr `.cat-right-ear`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `100px` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

Ihr `.cat-right-ear`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `white` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-left`-Eigenschaft mit dem Wert `35px solid blue` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-right`-Eigenschaft mit dem Wert `35px solid blue` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-top`-Eigenschaft mit dem Wert `35px solid red` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-bottom`-Eigenschaft mit dem Wert `35px solid red` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
