---
id: 646c48df8674cf2b91020ecc
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Sie haben einen grundlegenden Boilerplate-Code erhalten. Beginnen Sie damit, ein `link`-Element innerhalb Ihres `head`-Elements hinzuzufügen. Für dieses `link`-Element setzen Sie das `rel`-Attribut auf `stylesheet` und das `href`-Attribut auf `./styles.css`.

# --hints--

Sie sollten ein `link`-Element haben.

```js
assert.exists(document.querySelector("head>link"));
```

Ihr `link`-Element sollte ein `rel`- und ein `href`-Attribut haben, die auf `stylesheet` und `./styles.css` gesetzt sind.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
