---
id: 6476fb7cbaafa36d65e9cf35
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`fixed` ist ein `position`-Eigenschaftswert, mit dem Sie ein Element auf der Seite fixieren können, egal, wohin der Benutzer auf der Seite scrollt.

Sie müssen noch weitere Auszeichnungen vornehmen, um zu sehen, wie `fixed`-Positionierung funktioniert. Erstellen Sie in Ihrem HTML ein `div`-Element mit der Klasse `box`.

# --hints--

Sie sollten ein `div`-Element haben.

```js
assert.lengthOf(document.querySelectorAll('div'), 2);
```

Ihr `div`-Element sollte die Klasse `box` haben.

```js
assert.equal(document.querySelectorAll('div')[1]?.getAttribute('class'), 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
