---
id: 6476f5c17f99146071ee884c
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

Sie konnten sehen, dass nichts passiert. Das `.cat-head`-Element hat sich trotz der Einstellung von jeweils `top` und `left` von `100px` nicht bewegt. Aber das ist bei der `relative`-Positionierung nicht der Fall.

Wenn Sie den Wert `relative` verwenden, wird das Element weiterhin entsprechend dem normalen Fluss des Dokuments positioniert, aber die Werte `top`, `left`, `bottom` und `right` werden aktiv.

Geben Sie Ihrem `static` statt `.cat-head` eine Position von `relative` und lassen Sie sowohl die Eigenschaften `top` als auch `left` unverändert.

# --hints--

Ihr `.cat-head`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `relative` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Ihr `.cat-head` Selektor sollte eine `top` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Ihr `.cat-head` Selektor sollte eine `left` Eigenschaft mit dem Wert `100px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
