---
id: 6476f7a4827bcc61682f2347
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Wenn Sie den Wert `absolute` für Ihre `position`-Eigenschaft verwenden, wird das Element aus dem normalen Fluss des Dokuments genommen, und seine Position wird dann durch die Eigenschaften `top`, `right`, `bottom` und `left` bestimmt.

Setzen Sie die `position`-Eigenschaft Ihres `.cat-head`-Elements auf `absolute`.

# --hints--

Ihr `.cat-head`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `absolute` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'absolute');
```

Ihr `.cat-head`-Selektor sollte eine `top`-Eigenschaft mit einem beliebigen Wert in `px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/);
```

Ihr `.cat-head`-Selektor sollte eine `left`-Eigenschaft mit einem beliebigen Wert in `px` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>fCC Cat Painting</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class="cat-head"></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
--fcc-editable-region--
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
