---
id: 64770351e8586671ec0911f0
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

`sticky`-Positionierung ist eine Mischung aus `relative`- und `fixed`-Positionierung. Sie ermöglicht es einem Element, basierend auf der Scrollposition an einer bestimmten Stelle innerhalb seines enthaltenden Elements oder Viewports **kleben zu bleiben**.

Ändern Sie den Wert der `position`-Eigenschaft von `.cat-head` auf `sticky`, setzen Sie `top` auf `0` und entfernen Sie dann `left` und dessen Wert.

**Hinweis**: Um zu sehen, wie `sticky` funktioniert, müssen Sie einige Texte vor und nach Ihrem `.cat-head`-Element platzieren. Wenn Sie danach nach unten scrollen, sehen Sie, dass das `.cat-head` oben kleben bleibt und dort verweilt.

# --hints--

Ihr `.cat-head`-Selektor sollte eine `position`-Eigenschaft mit dem Wert `sticky` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'sticky');
```

Ihr `.cat-head`-Selektor sollte eine `top`-Eigenschaft mit dem Wert `0` haben. Achten Sie darauf, ein Semikolon hinzuzufügen.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '0px');
```

Sie sollten die `left`-Eigenschaft und deren Wert nicht in Ihrem Code haben.

```js
assert.notMatch(code, /left:\s*100px;?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
