---
id: 647d855ac12cd436059acd39
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Verwenden Sie den universellen Selektor, um `box-sizing: border-box;` zu Ihrem CSS hinzuzufügen. So stellen Sie sicher, dass Elemente Innenabstand und Rahmen in ihrer angegebenen Breite und Höhe berücksichtigen.

# --hints--

Sie sollten den universellen Selektor (`*`) haben.

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

Ihr universeller Selektor sollte eine `box-sizing`-Eigenschaft mit dem Wert `border-box` haben.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
