---
id: 6676a8b01e56ec1a1e07c202
title: Schritt 21
challengeType: 0
dashedName: step-21
---

# --description--

Beachten Sie, dass Sie jetzt ein weißes Quadrat mit dicken Rahmen haben. Die Rahmen haben diagonale Kanten, die verwendet werden können, um Dreiecke zu erstellen. Sie werden diese brauchen, um die Ohren Ihrer Katze zu gestalten.

Passen Sie innerhalb desselben Klassenselektors `height` und `width` zu `0` an. Setzen Sie den linken, rechten und oberen Rahmen auf `transparent`.

Entfernen Sie die `background-color`-Eigenschaft, und Sie sollten ein Dreieck sehen.

# --hints--


Sie sollten einen `.cat-right-ear`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Ihr `.cat-right-ear`-Selektor sollte eine `height`-Eigenschaft mit dem Wert `0` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '0px')
```

Ihr `.cat-right-ear`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `0` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '0px')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-left`-Eigenschaft mit dem Wert `35px solid transparent` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid transparent')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-right`-Eigenschaft mit dem Wert `35px solid transparent` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid transparent')
```

Ihr `.cat-right-ear`-Selektor sollte eine `border-top`-Eigenschaft mit dem Wert `35px solid transparent` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid transparent')
```


Ihr `.cat-right-ear`-Selektor sollte eine `border-bottom`-Eigenschaft mit dem Wert `35px solid red` haben. Vergessen Sie nicht, ein Semikolon hinzuzufügen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

Ihr `.cat-right-ear`-Selektor sollte keine `background-color`-Eigenschaft haben.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
