---
id: 6676a8a8d9861319ee901a90
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

Jetzt können Sie mit den Ohren Ihrer Katze beginnen.

Bereinigen Sie Ihren überprüften Code, indem Sie den `.cat-right-ear`-Selektor und alle seine Eigenschaften entfernen.

Verwenden Sie einen Klassenselektor, um dem `.cat-left-ear`-Element jeweils einen linken und rechten Rahmen von `35px solid transparent` zu geben. Setzen Sie außerdem den unteren Rahmen auf `70px solid #5e5e5e`.

# --hints--


Ihr `.cat-right-ear`-Selektor sollte keine `height`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
}
```

Ihr `.cat-right-ear`-Selektor sollte keine `width`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
}
```


Ihr `.cat-right-ear`-Selektor sollte keine `border-left`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
}
```

Ihr `.cat-right-ear`-Selektor sollte keine `border-right`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
}
```

Ihr `.cat-right-ear`-Selektor sollte keine `border-top`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
}
```

Ihr `.cat-right-ear`-Selektor sollte keine `border-bottom`-Eigenschaft haben.

```js
if (new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear')) {
  assert.isEmpty(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
}
```

Sie sollten keinen `.cat-right-ear`-Selektor haben.

```js
assert.notExists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Sie sollten einen `.cat-left-ear`-Selektor haben.

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-left-ear'))
```

Ihr `.cat-left-ear`-Selektor sollte eine `border-left`-Eigenschaft mit dem Wert `35px solid transparent` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft, '35px solid transparent')
```

Ihr `.cat-left-ear`-Selektor sollte eine `border-right`-Eigenschaft mit dem Wert `35px solid transparent` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight, '35px solid transparent')
```

Ihr `.cat-left-ear`-Selektor sollte eine `border-bottom`-Eigenschaft mit dem Wert `70px solid #5e5e5e` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom, '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
