---
id: 5dc174fcf86c76b9248c6eb2
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Los elementos HTML tienen etiquetas de apertura como `<h1>` y etiquetas de cierre como `</h1>`. El texto que mostrará un elemento va entre sus etiquetas de apertura y cierre.

Cambie el texto del elemento `h1` a continuación de `Hello World` a `CatPhotoApp` y observe el cambio en la vista previa del navegador.

Cuando hayas terminado, presiona el botón "Comprueba tu código" para ver si es correcto.

# --hints--

El texto `CatPhotoApp` debe estar presente en el código. Tal vez quieras revisar tu ortografía.

```js
assert(code.match(/catphotoapp/i));
```

Tu elemento `h1` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Tu elemento `h1` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del caracter `<`.

```js
assert(code.match(/<\/h1\>/));
```

Tienes más de un elemento `h1`. Elimina el elemento `h1` extra.

```js
assert(document.querySelectorAll('h1').length === 1);
```

El texto de tu elemento `h1` debe ser `CatPhotoApp`. O has omitido el texto, tienes un error tipográfico, o no está entre las etiquetas de apertura y cierre del elemento `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Parece que estás utilizando una extensión del navegador que está modificando la página. Asegúrate de desactivar todas las extensiones del navegador.

```js
if(__checkForBrowserExtensions){
  assert.isAtMost(document.querySelectorAll('script').length, 2);
  assert.equal(document.querySelectorAll('style').length, 1);
  assert.equal(document.querySelectorAll('link').length, 0);
}
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```

