---
id: 5dc24165f86c76b9248c6ebc
title: Paso 9
challengeType: 0
dashedName: step-9
---

# --description--

All `img` elements should have an `alt` attribute. The `alt` attribute's text is used for screen readers to improve accessibility and is displayed if the image fails to load.

Aquí hay un ejemplo de un elemento `img` con un atributo `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Inside the `img` element, add an `alt` attribute with this text:

`A cute orange cat lying on its back`

# --hints--

Tu código debe tener un elemento `img`. You removed the `img` element from an earlier step.

```js
assert(document.querySelector('img'));
```

Your `img` element does not have an `alt` attribute. Comprueba que hay un espacio después del nombre de la etiqueta de apertura y/o que hay espacios antes de todos los nombres de los atributos.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

Your `img` element's `alt` attribute value is set to something other than 'A cute orange cat lying on its back'. Make sure the `alt` attribute's value is surrounded with quotation marks.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```

