---
id: 5ddbd81294d8ddc1510a8e56
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

El texto de un enlace debe colocarse entre las etiquetas de apertura y cierre de un elemento de anclaje (`a`).

Aquí tienes un ejemplo de un enlace con el texto `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Añade el texto de ancla `cat photos` al elemento de ancla. Esto se convertirá en el texto del enlace.

# --hints--

Tu elemento anchor (`a`) debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Tu elemento anchor (`a`) debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert(code.match(/<\/a\>/));
```

El texto de tu elemento de ancla (`a`) debe ser `cat photos`. Asegúrate de poner el texto del enlace entre la etiqueta de apertura y etiqueta de cierre del elemento (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```

