---
id: 5ef9b03c81a63668521804dc
title: Step 46
challengeType: 0
dashedName: step-46
---

# --description--

Puedes usar botones de radio (radio button) para preguntas en las que quieres que el usuario solo te dé una respuesta a partir de múltiples opciones.

Here is an example of a radio button with the text set as `cat`:

```html
<input type="radio"> cat
```

Recuerda que un elemento de `input` es un elemento vacío.

Before the text input, add a radio button with the text set as:

`Indoor`

# --hints--

Debes crear un elemento `input` para tu botón de radio. Revisa la sintaxis.

```js
assert(document.querySelectorAll('form > input').length >= 2);
```

Tu elemento `input` debe tener una etiqueta de apertura, pero no una etiqueta de cierre.

```js
assert(document.querySelector('form > input') && !code.match(/<\/input\>/g));
```

Debes tener solo un elemento `input` agregado para tu botón de radio. Retire cualquier extra.

```js
assert(document.querySelectorAll('form > input'), 2);
```

Your new `input` element should be above the existing `input` with `type` attribute set to `text`. Los tienes en el orden incorrecto.

```js
const existingInputElem = document.querySelector('form > input[type="text"]');
assert(
  existingInputElem &&
    existingInputElem.previousElementSibling.nodeName === 'INPUT'
);
```

Tu nuevo elemento `input` no tiene un atributo `type`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Tu nuevo elemento `input` debe tener solo un atributo `type`. Retire cualquier extra.

```js
assert(document.querySelector('input')
  ?.getAttributeNames()
  .filter(attr => attr === 'type')
  .length === 1
);
```

Your new `input` element should have a `type` attribute with the value `radio`. Ha omitido el valor o tiene un error tipográfico. Recuerde que los valores de los atributos deben estar entre comillas.

```js
assert.match(
  document.querySelector('input')?.getAttribute('type'),
  /^radio$/i
);
```

Aunque ha establecido el atributo `type` del nuevo elemento `input` en `radio`, se recomienda encerrar siempre el valor de un atributo entre comillas.

```js
assert(!/\<\s*input\s+type\s*=\s*radio/i.test(code));
```

El texto `Indoor` del botón `radio` debe ubicarse después de él en lugar de antes.

```js
const radioInputElem = document.querySelector('input');
assert.notMatch(radioInputElem?.previousSibling?.nodeValue, /Indoor/i);
```

El texto `Indoor` debe ubicarse directamente a la derecha de su botón `radio`. Ha omitido el texto o tiene un error tipográfico.

```js
const radioInputElem = document.querySelector('input');
assert.match(radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '), /\s*Indoor/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--

          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

