---
id: 5ef9b03c81a63668521804e2
title: Paso 56
challengeType: 0
dashedName: step-56
---

# --description--

Los formularios Usualmente usan casillas de verificación (checkboxes-checkbox) para preguntas que puedan tener más de una respuesta. El elemento `input` con un atributo `type` establecido en `checkbox` crea una casilla de verificación.

Bajo el elemento `legend` que acabas de añadir, añade un `input` con su atributo `type` establecido a `checkbox` y su texto establecido a:

`Loving`

# --hints--

Tu elemento `input` de tu checkbox debe tener una etiqueta de apertura, pero no una etiquete de cierre.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Solo debiste de haber añadido un elemento input para tu checkbox. Elimina los que estén de más.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Tu nuevo elemento `input` debe estar debajo del elemento `legend` que tiene el texto `What's your cat's personality?`. Los tienes en el orden incorrecto.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Tu nuevo elemento `input` no tiene un atributo `type`. Comprueba que hay un espacio después del nombre de la etiqueta de apertura.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Tu nuevo elemento `input` debe tener un atributo `type` con el valor `checkbox`. Probablemente no has añadido el valor o tienes un error tipográfico. Recuerda que los valores de los atributos deben estar entre comillas.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Aunque al atributo `type` del elemento `input` le hayas dado el valor `checkbox`, se recomienda siempre poner el valor de un atributo entre comillas.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

El texto `Loving` debe estar inmediatamente a la derecha de tu checkbox. Asegúrate de que hay un espacio entre el elemento y el texto. Has omitido el texto o tienes un error tipográfico.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>

--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

