---
id: 671141d8e32fe934c26fa1be
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Añade las palabras `See more` antes del elemento ancla y `in our gallery` después del elemento de fondo.

# --hints--

Debe utilizar el elemento enlace existente, no crear uno nuevo.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

El texto dentro del elemento anclaje no debe cambiar, debe ser `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Debes tener `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` en tu código.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
