---
id: 671141feba228a35cefba82d
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Convierte el texto existente `cute cats` en un elemento de ancla que enlaza:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Debe haber un nuevo elemento de ancla en el primer elemento `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

El texto de ancla debe ser `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText, "cute cats");
```

El atributo `href` del nuevo elemento anclaje debe ser `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Solo debe haber una instancia de la frase `cute cats` en tu código.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

El texto del elemento `p` debe seguir siendo `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText, "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
