---
id: 66b205e6eacba4c4e54ea434
title: Erstellen Sie einen E-Mail-Maskierer
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

In diesem Labor maskieren Sie den Benutzernamen-Teil einer E-Mail-Adresse mit Sternchen. Maskieren ist ein Begriff, der verwendet wird, um sensible Informationen mit Sternchen oder anderen Zeichen zu verbergen oder zu ersetzen.

Wenn die E-Mail-Adresse zum Beispiel `myEmail@email.com` war, dann wird die maskierte E-Mail-Adresse `m*****l@email.com` sein.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Erstellen Sie eine Funktion mit dem Namen `maskEmail`, die `email` als Argument entgegennimmt.
2. Innerhalb der Funktion sollten Sie den `email` maskieren und den Domainnamen anhängen. Denken Sie daran, dass Sie Methoden wie `slice`, `repeat`, `indexOf` oder sogar `replace` verwenden können, um Ihnen zu helfen.
3. Außerhalb der Funktion deklarieren Sie eine Variable mit dem Namen `email`, um die E-Mail-Adresse zu speichern, die Sie maskieren möchten.
4. Rufen Sie die `maskEmail`-Funktion mit der `email`-Variablen auf und geben Sie das Ergebnis in der Konsole aus.
5. `maskEmail("apple.pie@example.com")` sollte `"a*******e@example.com"` zurückgeben.
6. `maskEmail("freecodecamp@example.com")` sollte `"f**********p@example.com"` zurückgeben.
7. `maskEmail("info@test.dev")` sollte `"i**o@test.dev"` zurückgeben.
8. `maskEmail("user@domain.org")` sollte `"u**r@domain.org"` zurückgeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `maskEmail` definieren.

```js
assert.isFunction(maskEmail);
```

Die `maskEmail`-Funktion sollte einen String, `email`, als Argument entgegennehmen.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

Außerhalb der Funktion sollten Sie eine `email`-Variable haben.

```js
assert.isDefined(email);
```

Sie sollten Ihrer `email`-Variablen eine gültige E-Mail-Adresse zuweisen.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")` sollte `"a*******e@example.com"` zurückgeben.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")` sollte `"f**********p@example.com"` zurückgeben.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")` sollte `"i**o@test.dev"` zurückgeben.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")` sollte `"u**r@domain.org"` zurückgeben.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

Ihre `maskEmail` sollte das korrekte Ergebnis liefern.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

Sie sollten die Ausgabe des Aufrufs von `maskEmail` mit `email` als Argument in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
