---
id: 680a3d4c065557260de2ff33
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذه الورشة، ستنشئ برنامجًا يسمى caesar cipher وهو طريقة تشفير تُحوّل الحروف في الأبجدية لتحويل الرسائل إلى رموز.

ابدأ بإنشاء متغير يسمى `shift` وقم بتعيين القيمة `5` إلى متغيرك الجديد.

# --hints--

يجب إعلان متغير يسمى `shift`. انتبه لوضع اسم المتغير في بداية السطر.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

يجب تعيين القيمة `5` إلى متغير `shift` الخاص بك.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
