---
id: 680a7e200400089967b8b1f4
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

تأخذ الطريقة `str.maketrans()` سلسلتين نصيتين متساويتين في الطول وتُرجع جدول ترجمة يربط كل حرف من السلسلة الأولى بالحرف المقابل له في السلسلة الثانية. يُخزن كل حرف في جدول الترجمة كرقم ترتيبي في Unicode، وهو رقم يميز الحرف بشكل فريد.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

سوف تتعلم المزيد عن نوع البنية التي تُرجعها `str.maketrans()` لاحقًا في المنهج. الآن، أنشئ جدول ترجمة يربط الأحرف في `alphabet` بالأحرف في `shifted_alphabet` وعيّنه في متغير باسم `translation_table`.

# --hints--

يجب إعلان متغير باسم `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

يجب استدعاء `str.maketrans()` مع تمرير `alphabet` و `shifted_alphabet` كمعلمات، وتعيين النتيجة إلى `translation_table`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
