---
id: 680a836235461fce1ff53b3f
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

تأخذ الطريقة `translate()` المعلمة وهي جدول الترجمة الذي يُولد بواسطة `str.maketrans()`. تُستدعى على سلسلة نصية وتُعيد نسخة من السلسلة الأصلية حيث تم استبدال الأحرف بناءً على جدول الترجمة:

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

استدعِ الطريقة `translate()` على `text` مع تمرير `translation_table` كمعلمة، وعيّن النتيجة إلى متغير باسم `encrypted_text`.

# --hints--

يجب أن يكون لديك متغير باسم `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

يجب أن تعيّن `text.translate(translation_table)` إلى `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
