---
id: 680a431692a5f863f726c438
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في هذه الورشة، ستبني *شيفرة قيصر*. هذه واحدة من أبسط التقنيات لتشفير النص، وتعتمد على استبدال كل حرف من النص الأصلي بالحرف الموجود بعدد ثابت من المواقع في الأبجدية. على سبيل المثال، مع إزاحة مقدارها 5، يُستبدل `a` بـ `f`، و`b` بـ `g` وهكذا.

لتنفيذ هذه الشيفرة، تحتاج إلى إنشاء نسخة جديدة من الأبجدية تبدأ عند الموضع المحدد بواسطة الإزاحة. كما تعلمت في درس سابق، يمكنك استخراج جزء من سلسلة نصية باستخدام تقطيع السلاسل النصية:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

أنشئ متغيرًا باسم `shifted_alphabet` واستخدم صيغة التقطيع لتعيين الجزء من `alphabet` الذي يبدأ عند فهرس `shift`. ثم، استدعِ الدالة المدمجة `print()` لطباعة `shifted_alphabet` على الطرفية واطلع على النتيجة.

# --hints--

يجب إعلان متغير جديد باسم `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

يجب تعيين `alphabet[shift:]` إلى `shifted_alphabet`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

يجب استدعاء الدالة `print()` مع تمرير `shifted_alphabet` كمعلمة.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
