---
id: 680a41ef5bfdc059cfd9d4b5
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

أعلن عن متغير آخر اسمه `alphabet` وقم بتعيين السلسلة النصية `abcdefghijklmnopqrstuvwxyz` لهذا المتغير.

# --hints--

يجب أن تعلن عن متغير اسمه `alphabet`. انتبه لوضع اسم المتغير في بداية السطر.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

يجب أن تعيّن السلسلة النصية `abcdefghijklmnopqrstuvwxyz` إلى متغير `alphabet`. تذكر استخدام علامات اقتباس مفردة أو علامات اقتباس مزدوجة لتضمين السلسلة النصية.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
