---
id: 680a549f60f5a3fedfc0edd9
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

كما ترى من المخرجات، يبدأ الأبجدية المنقولة بالحرف `f` لأن قيمة `shift` هي `5`. لكن الآن الحروف الخمسة الأولى من الأبجدية – `a`، `b`، `c`، `d` و `e` – مفقودة من الأبجدية المنقولة، لذا ستحتاج إلى إضافتها في نهاية الأبجدية المنقولة. يمكنك استخدام العامل `+` لربط الجزء المفقود من الأبجدية.

عدّل التعيين الحالي للمتغير `shifted_alphabet`: استخدم صيغة التقطيع لاستخراج الجزء الأول المفقود من `alphabet` واربطه بـ `alphabet[shift:]`.

كتذكير، `sentence[start:stop]` تُرجع الأحرف من `sentence` من الموضع `start` (شامل) إلى `stop` (غير شامل).

# --hints--

يجب ألا تعدل المتغير `alphabet` مباشرة.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

يجب أن تعيّن ربط `alphabet[shift:]` والجزء الأول المفقود من `alphabet` إلى المتغير `shifted_alphabet` في نفس السطر الذي أُعلن فيه. استخدم `shift` لتحديد نقطة التوقف في التقطيع.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
