---
id: 680a82806c4c6bc3567f2708
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

أعلن عن متغير جديد باسم `text` وقم بتعيين السلسلة النصية `hello world` له. ستكون هذه الرسالة التي ستُشفّر.

# --hints--

يجب إعلان متغير باسم `text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

يجب تعيين السلسلة النصية `hello world` إلى `text`. تذكر أن تحيط السلسلة النصية إما بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة واهتم بحالة الحروف.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
