---
id: 6810891bd6de8a87a833df42
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

كما ترى من الناتج، تم تشفير الرسالة. الخطوة التالية ستكون جعل شفرتك قابلة لإعادة الاستخدام في حال أردت تشفير رسائل مختلفة.

لهذا، ستحتاج إلى إنشاء دالة. كتذكير، إليك كيفية إنشاء دالة باسم `spam` تقوم بطباعة `Spam!` على الطرفية:

```py
def spam():
    print('Spam!')
```

أنشئ دالة باسم `caesar`. ضع كل شفرتك الحالية داخل جسم الدالة. انتبه للحفاظ على نفس مستوى التراجع (المسافة البادئة) لكل الأسطر داخل جسم الدالة.

# --hints--

يجب أن يكون لديك دالة باسم `caesar`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

يجب أن تنقل كل الشفرة التي كتبتها حتى الآن داخل جسم دالة `caesar`. تأكد من الحفاظ على نفس مستوى التراجع لكل الأسطر داخل جسم الدالة...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
