---
id: 6818b80519e6c6c47a8bb6e8
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

كما تعلمت في درس سابق، يمكن أن تحتوي الدوال على معلمات، وهي متغيرات يمكن الرجوع إليها داخل الدالة. هذه دالة `sum` تحتوي على معلمتين، `num1` و `num2`.

```py
def sum(num1, num2):
    print(num1 + num2)
```

الرسالة التي يجب تشفيرها وقيمة الإزاحة ما زالتا مكتوبتين مباشرة في دالتك. أعطِ دالتك معلمتين باسمَي `text` و `shift`، واحذف إعلانات متغيري `text` و `shift` من جسم دالة `caesar`.

# --hints--

يجب أن تحتوي دالة `caesar` على معلمتين، `text` و `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

يجب ألا يكون لديك متغير `text` معلن داخل جسم دالة `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

يجب ألا يكون لديك متغير `shift` معلن داخل جسم دالة `caesar`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
