---
id: 6818ed1adc249bb91d373729
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

في الدروس السابقة، تعلمت أن جملة `return` تُستخدم لإرجاع قيمة من دالة، بحيث يمكنك استخدامها في مكان آخر في شفرتك:

```py
def spam():
    return 'Spam!'
```

احذف استدعاء `print(encrypted_text)` من دالتك. ثم احذف إعلان المتغير `encrypted_text` وأرجع مباشرةً `text.translate(translation_table)` من دالتك.

# --hints--

يجب ألا يحتوي دالة `caesar` على استدعاء `print(encrypted_text)`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

يجب ألا يحتوي دالة `caesar` على متغير `encrypted_text`.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

يجب أن تُرجع دالة `caesar` القيمة `text.translate(translation_table)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
