---
id: 6818e9c7c8dc5694e8f64fc3
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

اختبر الآن الدالة `caesar` باستدعائها مع السلسلة النصية `freeCodeCamp` والعدد `3` كمعلمتين. عيّن استدعاء الدالة إلى متغير باسم `encrypted_text`.

# --hints--

يجب أن يكون لديك متغير باسم `encrypted_text` في النطاق العام.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

يجب تعيين `caesar('freeCodeCamp', 3)` إلى المتغير `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
