---
id: 6818fcb250f34e62cab32c39
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

ربما لاحظت أنه، رغم تشفير الرسالة، لم تُلمس الأحرف الكبيرة. يحدث هذا لأن جدول الترجمة لا يشمل الأحرف الكبيرة.

لإصلاح ذلك، عليك تعديل السلاسل النصية الممررة إلى استدعاء `str.maketrans()` باستخدام الطريقة `upper()`:

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

حدّث استدعاء `str.maketrans()` بضم النسخة الكبيرة من كل معلمة إلى المعلمة نفسها.

# --hints--

يجب تعديل المعلمات في استدعاء `str.maketrans()` بضم النسخة الكبيرة من كل معلمة إلى المعلمة نفسها.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
