---
id: 6818eb3f98c0bca5af7c8d03
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

الآن شفرتك قابلة لإعادة الاستخدام. مع ذلك، دالة `caesar` تطبع رسالة على الطرفية وتُرجع `None` افتراضيًا. لإثبات ذلك، اطبع `encrypted_text` في نهاية شفرتك.

# --hints--

يجب أن تطبع `encrypted_text` خارج جسم الدالة.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
