---
id: 6819bc0637b80256a33adccc
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

الآن بعد أن نفذت الوظائف الأساسية للتشفير، حان الوقت لإضافة بعض التحقق. لهذا، ستحتاج إلى جملة `if`.

في بداية جسم الدالة، أنشئ جملة `if`. استخدم `True` كشرط مؤقتًا، وداخل جسم جملة `if` أرجع السلسلة النصية `Shift must be an integer value.`

# --hints--

يجب أن تحتوي دالتك `caesar` على جملة `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

يجب أن تستخدم جملة `if` شرط `True`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

يجب أن تُرجع من جملة `if` السلسلة النصية `Shift must be an integer value.`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
