---
id: 6819ca464db219ab86f37dc4
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

كما تحققت سابقًا، يجب أن يكون الإزاحة الممررة لتشفير النص عددًا صحيحًا موجبًا. مع ذلك، لا يمكن أن تتجاوز `25` (آخر فهرس في `alphabet`).

أضف شرطًا ثانيًا إلى جملة `if` التي تتحقق مما إذا كانت قيمة `shift` أكبر من `25`. تذكر أن عملية OR المنطقية في بايثون تُنفذ عبر العامل `or`.

أيضًا، حدّث الرسالة المعادة إلى `Shift must be an integer between 1 and 25.`

# --hints--

يجب أن تستخدم العامل `or` لإضافة شرط ثاني بعد `shift < 1` يتحقق مما إذا كانت قيمة `shift` أكبر من `25`. لا تعدل الشرط الموجود.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

يجب أن تعيد جملة `if` الثانية الرسالة `Shift must be an integer between 1 and 25.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
