---
id: 6819c7beb4a8289c2b4f24d5
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

يجب ألا تقبل الدالة أي قيمة إزاحة سالبة أو صفرية. لذلك، بعد جملة `if` الأولى، أنشئ جملة `if` أخرى تتحقق مما إذا كانت قيمة `shift` أقل من `1` وتعيد السلسلة النصية `Shift must be a positive integer.`

# --hints--

يجب أن تحتوي على جملة `if` ثانية تتحقق مما إذا كانت قيمة `shift` أقل من `1`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

يجب أن تُرجع جملة `if` الثانية السلسلة النصية `Shift must be a positive integer.`

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
