---
id: 6819cf38880d8cc9b9f6af7a
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

تسمح بايثون بتحديد قيمة افتراضية للمعلمات في الدالة، مما ينشئ دالة يمكن استدعاؤها بعدد أقل من المعلمات أو بدون معلمات. إليك كيفية إنشاء دالة بمعلمة `name` لها قيمة افتراضية:

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

أضف معلمة ثالثة باسم `encrypt` إلى دالتك وأعطها قيمة افتراضية `True`.

# --hints--

يجب أن تضيف معلمة ثالثة باسم `encrypt` بقيمة افتراضية `True` إلى دالتك. لا تعدل المعلمات الموجودة.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
