---
id: 681a44f4c3235f7d8f428545
title: الخطوة 25
challengeType: 20
dashedName: step-25
---

# --description--

الآن ستختبر الدالة `decrypt`. استبدل القيمة المُعيَّنة إلى `encrypted_text` بالسلسلة النصية التالية، التي تمثل رسالة لفك تشفيرها: `Pbhentr vf sbhaq va hayvxryl cynprf.`

بعد ذلك، أعلن عن متغير باسم `decrypted_text` وعيّن له استدعاء الدالة `decrypt` مع `encrypted_text` كمعامل أول وتحويل بمقدار `13` كمعامل ثاني.

أخيرًا، اطبع `decrypted_text` على الطرفية. بذلك، يكتمل تشفير قيصر.

# --hints--

يجب أن تعيّن `Pbhentr vf sbhaq va hayvxryl cynprf.` إلى `encrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'")`)) })
```

يجب أن يكون لديك متغير باسم `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("decrypted_text")`)) })
```

يجب أن تعيّن `decrypt(encrypted_text, 13)` إلى `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("decrypted_text").is_equivalent("decrypted_text = decrypt(encrypted_text, 13)")`)) })
```

يجب أن تطبع `decrypted_text`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(decrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = encrypt('freeCodeCamp', 3)
print(encrypted_text)
--fcc-editable-region--
```

# --solutions--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

encrypted_text = 'Pbhentr vf sbhaq va hayvxryl cynprf.'
decrypted_text = decrypt(encrypted_text, 13)
print(decrypted_text)
```
