---
id: 681a1d60e037127e7c8e0230
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

ستستخدم المعلمة التي أُضيفت في الخطوة السابقة لتحديد ما إذا كانت الدالة يجب أن تُشفّر النص الممرر إليها (السلوك الافتراضي، `encrypt=True`)، أو إذا كان يجب أن تفك تشفير رسالة مشفرة.

أنشئ جملة `if` تتحقق مما إذا كانت `encrypt` غير صحيحة. داخل جملة `if` الجديدة، عرّف `shift` لتساوي `- shift`. هذا ضروري لتمكين الإزاحة في الاتجاه المعاكس بالنسبة لعملية التشفير.

# --hints--

يجب أن تعرّف `shift` لتساوي `- shift` عندما تكون `encrypt` غير صحيحة.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
