---
id: 681a2a6cf95e82ca30866364
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

أعلن دالتين باسم `encrypt` و `decrypt`، كلتاهما تحتويان على معلمتين `text` و `shift`.

استخدم `encrypt` لعملية التشفير، و `decrypt` لفك التشفير، مع تسمية كلتا العمليتين بمسمى وصفي.

أرجع استدعاء `caesar` مع تمرير `text` و `shift` من كلا الدالتين الجديدتين، لكن تأكد من تمرير `False` كالمعلمة الثالثة في دالة `decrypt`.

# --hints--

يجب أن يكون لديك دالة باسم `encrypt` تحتوي على معلمتين، `text` و `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

يجب أن تُرجع دالة `encrypt` استدعاء `caesar(text, shift)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

يجب أن يكون لديك دالة باسم `decrypt` تحتوي على معلمتين، `text` و `shift`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

يجب أن تُرجع دالة `decrypt` استدعاء `caesar(text, shift, False)`.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
