---
id: 673b3d6b7ef7318eef926d5a
title: Crea una mood board
challengeType: 25
dashedName: build-a-mood-board
demoType: onClick
---

# --description--

Una mood board è un collage di immagini e testo che trasmette un’idea generale, un obiettivo o una sensazione su un argomento specifico.

In questo laboratorio, creerai una mood board usando componenti React riutilizzabili. Il CSS ti è già stato fornito.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare ed esportare un componente `MoodBoardItem` che accetta tre props: `color`, `image` e `description`.
2. Il tuo componente `MoodBoardItem` dovrebbe restituire un `div` con la classe `mood-board-item` come elemento di livello superiore.
3. Dovresti impostare il colore di sfondo dell’elemento `.mood-board-item` al valore della prop `color` usando stili in linea.
4. Dovresti renderizzare un elemento `img`, con classe `mood-board-image` e il suo attributo `src` impostato al valore della prop `image`, all’interno dell’elemento `.mood-board-item`.
5. Dovresti renderizzare un elemento `h3`, con classe `mood-board-text` e il suo testo impostato al valore della prop `description`, all’interno dell’elemento `.mood-board-item`.
6. Dovresti creare ed esportare un `MoodBoard`.
7. Il tuo componente `MoodBoard` dovrebbe restituire un `div` come elemento di livello superiore.
8. Il tuo componente `MoodBoard` dovrebbe renderizzare un elemento `h1` con classe `mood-board-heading` e il testo `Destination Mood Board`.
9. Il tuo componente `MoodBoard` dovrebbe renderizzare un `div` con classe `mood-board`.
10. Il tuo componente `MoodBoard` dovrebbe renderizzare almeno tre componenti `MoodBoardItem` all’interno dell’elemento `.mood-board`, ognuno dovrebbe passare le props `color`, `image` e `description` con valori validi.

Puoi usare le seguenti immagini nella tua mood board se vuoi:

- `https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/shore.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/grass.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/ship.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg`
- `https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg`

# --hints--

Dovresti esportare un componente `MoodBoardItem`.

```js
assert.isFunction(window.index.MoodBoardItem, 1);
```

Il tuo componente `MoodBoardItem` dovrebbe restituire un `div` con classe `mood-board-item` come elemento di livello superiore.

```js
const item = document.getElementsByClassName('mood-board-item')[0];
assert.exists(item);
assert.equal(item.tagName, 'DIV');
```

Il colore di sfondo dell’elemento `.mood-board-item` dovrebbe essere impostato al valore della prop `color` usando stili in linea.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { color: "red" });
  const moodBoardItem = container.querySelector(".mood-board-item");
  assert.equal(moodBoardItem.style.backgroundColor, "red");
```

Il tuo componente `MoodBoardItem` dovrebbe renderizzare un elemento `img` con classe `mood-board-image` e il suo `src` impostato al valore della prop `image`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { image: "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg" });
  const img = container.querySelector(".mood-board-image");
  assert.exists(img);
  assert.equal(img.tagName, "IMG");
  assert.equal(img.src, "https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg");
```

Il tuo componente `MoodBoardItem` dovrebbe renderizzare un elemento `h3` con classe `mood-board-text` e il suo testo impostato al valore della prop `description`.

```js
  const container = await __helpers.prepTestComponent(window.index.MoodBoardItem, { description: "Caribbean" });
  const text = container.querySelector(".mood-board-text");
  assert.exists(text);
  assert.equal(text.tagName, "H3");
  assert.equal(text.textContent, "Caribbean");
```

Dovresti esportare un componente `MoodBoard`.

```js
assert.isFunction(window.index.MoodBoard, 1);
```

Il tuo componente `MoodBoard` dovrebbe restituire un `div` come elemento di livello superiore.

```js
const root = document.getElementById('root');
assert.equal(root.children[0].tagName, 'DIV');
```

Il tuo componente `MoodBoard` dovrebbe renderizzare un elemento `h1` con classe `mood-board-heading` e il testo `Destination Mood Board`.

```js
const heading = document.getElementsByClassName('mood-board-heading')[0];
assert.exists(heading);
assert.equal(heading.tagName, 'H1');
assert.equal(heading.textContent, 'Destination Mood Board');
```

Il tuo componente `MoodBoard` dovrebbe renderizzare almeno tre componenti `MoodBoardItem`, ognuno dovrebbe passare le props `color`, `image` e `description` con valori validi.

```js
const moodboard = document.getElementsByClassName('mood-board')[0];
const items = moodboard.querySelectorAll(".mood-board-item");
assert.isAtLeast(items.length, 3);     

items.forEach((item) => {
  const color = item.style.backgroundColor;
  const img = item.querySelector("img");
  const text = item.querySelector("h3");
  
  assert.isAtLeast(color.length, 1);
  assert.isAtLeast(img.src.length, 1);
  assert.isAtLeast(text.textContent.length, 1);
});
```

Il tuo componente `MoodBoard` dovrebbe essere renderizzato nell’elemento `#root` della pagina.

```js
const root = document.getElementById('root');
const moodBoard = root.getElementsByClassName('mood-board')[0];
assert.exists(moodBoard);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx

```

# --solutions--

```html
<!doctype html>
<html>
  <head>
    <meta charset="UTF-8" />
    <title>Mood Board</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.3/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
  </head>

  <body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { MoodBoard } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(
        <MoodBoard />
      );
    </script>
  </body>
</html>
```

```css
body {
  background-color: #ffffcc;
}

.mood-board-heading {
  text-align: center;
  font-size: 2.5em;
  color: #333;
  margin-top: 20px;
}

.mood-board {
  display: grid;
  grid-template-columns: repeat(3, 1fr);
  gap: 20px;
  padding: 20px;
  max-width: 900px;
  margin: 0 auto;
}

.mood-board-item {
  border-radius: 10px;
  padding: 10px;
  display: flex;
  flex-direction: column;
  align-items: center;
  justify-content: center;
  color: #fff;
  box-shadow: 0 4px 8px rgba(0, 0, 0, 0.2);
  text-align: center;
  height: 250px;
}

.mood-board-image {
  border-radius: 5px;
  width: 180px;
  height: 150px;
  object-fit: cover;
}

.mood-board-text {
  margin-top: 20px;
  font-size: 1.2em;
}
```

```jsx
export function MoodBoardItem(props) {
  return (
    <div className='mood-board-item' style={{ backgroundColor: props.color }}>
      <img src={props.image} alt='Mood' className='mood-board-image' />
      <h3 className='mood-board-text'>{props.description}</h3>
    </div>
  );
}

export function MoodBoard() {
  return (
    <div>
      <h1 className='mood-board-heading'>Destination Mood Board</h1>
      <div className='mood-board'>
        <MoodBoardItem
          color='#2da64f'
          image='https://cdn.freecodecamp.org/curriculum/labs/pathway.jpg'
          description='Caribbean'
        />
        <MoodBoardItem
          color='#8e44ad'
          image='https://cdn.freecodecamp.org/curriculum/labs/shore.jpg'
          description='Gawadar Beach'
        />
        <MoodBoardItem
          color='#3498db'
          image='https://cdn.freecodecamp.org/curriculum/labs/grass.jpg'
          description='Cape Town'
        />
        <MoodBoardItem
          color='#bf3d7e'
          image='https://cdn.freecodecamp.org/curriculum/labs/ship.jpg'
          description='Suez Canal'
        />
        <MoodBoardItem
          color='#e74c3c'
          image='https://cdn.freecodecamp.org/curriculum/labs/santorini.jpg'
          description='Santorini'
        />
        <MoodBoardItem
          color='#95a5a6'
          image='https://cdn.freecodecamp.org/curriculum/labs/pigeon.jpg'
          description='Istanbul'
        />
      </div>
    </div>
  );
}
```
