---
id: 56105e7b514f539506016a5e
title: Лічити назад за допомогою циклу for
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

Цикл for також може лічити назад, якщо ми правильно визначимо умови.

Щоб зменшувати значення на два під час кожної ітерації, нам потрібно змінити ініціалізацію, умову та фінальний вираз.

Ми почнемо з `i = 10` і будемо виконувати цикл, поки `i > 0`. Кожного разу зменшуватимемо `i` на 2 за допомогою `i -= 2`.

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray` тепер міститиме `[10, 8, 6, 4, 2]`. Змінимо ініціалізацію та фінальний вираз, щоб лічити назад по два і створити масив спадних непарних чисел.

# --instructions--

Додайте непарні числа від 9 до 1 у `myArray` за допомогою циклу `for`.

# --hints--

Вам слід використовувати цикл `for` для цього.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

Вам слід використовувати метод масиву `push`.

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray` має дорівнювати `[9, 7, 5, 3, 1]`.

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
