---
id: 56104e9e514f539506016a5c
title: Ітерувати непарні числа за допомогою циклу for
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

Цикли for не обов’язково ітерують по одному. Змінивши наш `final-expression`, ми можемо рахувати парними числами.

Почнемо з `i = 0` і будемо виконувати цикл, поки `i < 10`. Кожного разу збільшуватимемо `i` на 2 за допомогою `i += 2`.

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray` тепер міститиме `[0, 2, 4, 6, 8]`. Змінимо наш `initialization`, щоб рахувати непарні числа.

# --instructions--

Додайте непарні числа від 1 до 9 у `myArray` за допомогою циклу `for`.

# --hints--

Для цього слід використовувати цикл `for`.

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray` має дорівнювати `[1, 3, 5, 7, 9]`.

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
