---
id: 56533eb9ac21ba0edf2244b1
title: Складене присвоєння зі збільшеним множенням
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

Оператор `*=` помножує змінну на число.

```js
myVar = myVar * 5;
```

помножить `myVar` на `5`. Це можна переписати так:

```js
myVar *= 5;
```

# --instructions--

Перетворіть присвоєння для `a`, `b` і `c` так, щоб вони використовували оператор `*=`.

# --hints--

`a` має дорівнювати `25`.

```js
assert(a === 25);
```

`b` має дорівнювати `36`.

```js
assert(b === 36);
```

`c` має дорівнювати `46`.

```js
assert(c === 46);
```

Ви маєте використати оператор `*=` для кожної змінної.

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

Не змінюйте код вище за вказаний коментар.

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
