---
id: 56533eb9ac21ba0edf2244d8
title: Порівняння з логічним оператором and
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

Іноді потрібно перевірити більше ніж одну умову одночасно. <dfn>Логічний оператор and</dfn> (`&&`) повертає `true` тоді і тільки тоді, коли <dfn>операнди</dfn> зліва і справа від нього є істинними.

Такий самий ефект можна досягти, вклавши одну інструкцію `if` всередину іншої інструкції `if`.

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

Цей код поверне `Yes`, якщо `num` більше за `5` і менше за `10`. Ту саму логіку можна записати за допомогою <dfn>логічного оператора and</dfn>.

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

Замініть дві інструкції if однією інструкцією, використовуючи оператор `&&`, який поверне рядок `Yes`, якщо `val` менше або дорівнює `50` і більше або дорівнює `25`. Інакше поверне рядок `No`.

# --hints--

Ви повинні використати оператор `&&` один раз

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

У вас має бути лише одна інструкція `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)` має повертати рядок `No`

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)` має повертати рядок `No`

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)` має повертати рядок `Yes`

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)` має повертати рядок `Yes`

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)` має повертати рядок `Yes`

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)` має повертати рядок `No`

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)` має повертати рядок `No`

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)` має повертати рядок `No`

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
