---
id: 56533eb9ac21ba0edf2244d9
title: Порівняння з логічним оператором або
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

<dfn>Логічний оператор або</dfn> (`||`) повертає `true`, якщо будь-який із <dfn>операндів</dfn> є `true`. Інакше він повертає `false`.

<dfn>Логічний оператор або</dfn> складається з двох символів вертикальної риски: (`||`). Зазвичай його можна знайти між клавішами Backspace та Enter.

Наведену нижче конструкцію ви, ймовірно, вже бачили на попередніх етапах.

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

Цей код поверне `Yes`, якщо `num` знаходиться між `5` та `10` (включно з `5` та `10`). Ту ж логіку можна записати за допомогою <dfn>логічного оператора або</dfn>.

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

Об’єднайте дві `if` інструкції в одну, яка повертає рядок `Outside`, якщо `val` не знаходиться між `10` та `20` включно. Інакше поверніть рядок `Inside`.

# --hints--

Ви повинні використати оператор `||` один раз

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

У вас має бути лише одна `if` інструкція

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` має повертати рядок `Outside`

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` має повертати рядок `Outside`

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` має повертати рядок `Inside`

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` має повертати рядок `Inside`

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` має повертати рядок `Inside`

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` має повертати рядок `Inside`

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` має повертати рядок `Outside`

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` має повертати рядок `Outside`

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
