---
id: 56533eb9ac21ba0edf2244b5
title: Екранування буквальних лапок у рядках
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

Коли ви визначаєте рядок, потрібно починати і закінчувати його одинарною або подвійною лапкою. Що робити, якщо вам потрібна буквальна лапка: `"` або `'` всередині рядка?

У JavaScript ви можете <dfn>екранувати</dfn> лапку, щоб вона не сприймалась як кінець рядка, поставивши перед нею <dfn>зворотну скісну риску</dfn> (`\`).

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

Це сигналізує JavaScript, що наступна лапка не є кінцем рядка, а має з’явитися всередині рядка. Тож, якщо ви виведете це в консоль, отримаєте:

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

Використайте <dfn>зворотні скісні риски</dfn>, щоб призначити рядок змінній `myStr`, так щоб при виведенні в консоль ви побачили:

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

Вам слід використати дві подвійні лапки (`"`) і чотири екрановані подвійні лапки (`\"`).

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

Змінна `myStr` має містити рядок: `I am a "double quoted" string inside "double quotes".`

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
