---
id: 56533eb9ac21ba0edf2244cb
title: Маніпулювання складними об’єктами
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

Іноді ви можете захотіти зберігати дані у гнучкій <dfn>структурі даних</dfn>. Об’єкт JavaScript — це один зі способів працювати з гнучкими даними. Вони дозволяють довільно комбінувати <dfn>рядки</dfn>, <dfn>числа</dfn>, <dfn>булеві</dfn> значення, <dfn>масиви</dfn>, <dfn>функції</dfn> та <dfn>об’єкти</dfn>.

Ось приклад складної структури даних:

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

Це масив, який містить один об’єкт всередині. Об’єкт має різні метадані про альбом. Також він містить вкладений `formats` масив. Якщо ви хочете додати більше записів про альбоми, це можна зробити, додавши записи до верхнього рівня масиву. Об’єкти зберігають дані у властивостях, які мають формат ключ-значення. У наведеному вище прикладі `"artist": "Daft Punk"` — це властивість з ключем `artist` і значенням `Daft Punk`.

**Примітка:** Потрібно ставити кому після кожного об’єкта в масиві, якщо це не останній об’єкт у масиві.

# --instructions--

Додайте новий альбом до масиву `myMusic`. Додайте рядки `artist` і `title`, число `release_year` і масив рядків `formats`.

# --hints--

`myMusic` має бути масивом

```js
assert(Array.isArray(myMusic));
```

`myMusic` має містити щонайменше два елементи

```js
assert(myMusic.length > 1);
```

Елементи в масиві `myMusic` мають бути об’єктами

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

Ваш об’єкт у `myMusic` має містити щонайменше 4 властивості

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

Ваш об’єкт у `myMusic` має містити властивість `artist`, яка є рядком

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

Ваш об’єкт у `myMusic` має містити властивість `title`, яка є рядком

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

Ваш об’єкт у `myMusic` має містити властивість `release_year`, яка є числом

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

Ваш об’єкт у `myMusic` має містити властивість `formats`, яка є масивом

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats` має бути масивом рядків з щонайменше двома елементами

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
