---
id: 56533eb9ac21ba0edf2244cf
title: колекція записів
challengeType: 1
forumTopicId: 18261
dashedName: record-collection
---

# --description--

Ви створюєте функцію, яка допомагає в обслуговуванні колекції музичних альбомів. Колекція організована як об’єкт, що містить кілька альбомів, які також є об’єктами. Кожен альбом у колекції представлений унікальним `id` як назва властивості. Усередині кожного об’єкта альбому є різні властивості, що описують інформацію про альбом. Не всі альбоми мають повну інформацію.

Функція `updateRecords` приймає 4 аргументи, представлені такими параметрами функції:

-   `records` — об’єкт, що містить кілька окремих альбомів
-   `id` — число, що представляє конкретний альбом в об’єкті `records`
-   `prop` — рядок, що представляє назву властивості альбому для оновлення
-   `value` — рядок, що містить інформацію для оновлення властивості альбому

Завершіть функцію, використовуючи наведені нижче правила для зміни об’єкта, переданого у функцію.

-   Ваша функція завжди має повертати весь об’єкт `records`.
-   Якщо `value` є порожнім рядком, видаліть задану властивість `prop` з альбому.
-   Якщо `prop` не є `tracks` і `value` не є порожнім рядком, призначте `value` цій властивості `prop` альбому.
-   Якщо `prop` є `tracks` і `value` не є порожнім рядком, але альбом не має властивості `tracks`, створіть порожній масив і додайте до нього `value`.
-   Якщо проп є `tracks` і `value` не є порожнім рядком, додайте `value` у кінець існуючого масиву `tracks` альбому.

**Примітка:** Для тестів використовується копія об’єкта `recordCollection`. Ви не повинні безпосередньо змінювати об’єкт `recordCollection`.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Після `updateRecords(recordCollection, 5439, "artist", "ABBA")` `artist` має бути рядком `ABBA`

```js
assert(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'] ===
    'ABBA'
);
```

Після `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")` `tracks` має мати рядок `Take a Chance on Me` як останній і єдиний елемент.

```js
assert(
  updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me') &&
  _recordCollection[5439]['tracks'].length === 1 &&
  _recordCollection[5439]['tracks'].pop() === 'Take a Chance on Me'
);
```

Після `updateRecords(recordCollection, 2548, "artist", "")` `artist` не має бути встановлено

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert(!_recordCollection[2548].hasOwnProperty('artist'));
```

Після `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")` `tracks` має мати рядок `Addicted to Love` як останній елемент.

```js
assert(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop() === 'Addicted to Love'
);
```

Після `updateRecords(recordCollection, 2468, "tracks", "Free")` `tracks` має мати рядок `1999` як перший елемент.

```js
assert(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0] === '1999'
);
```

Після `updateRecords(recordCollection, 2548, "tracks", "")` `tracks` не має бути встановлено

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert(!_recordCollection[2548].hasOwnProperty('tracks'));
```

Після `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")` `albumTitle` має бути рядком `Riptide`

```js
assert(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ] === 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
// Setup
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  return records;
}

updateRecords(recordCollection, 5439, 'artist', 'ABBA');
```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

// Only change code below this line
function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
