---
id: 56533eb9ac21ba0edf2244c8
title: Доступ до властивостей об’єкта за допомогою дужкової нотації
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

Другий спосіб доступу до властивостей об’єкта — це дужкова нотація (`[]`). Якщо властивість об’єкта, до якої ви намагаєтеся отримати доступ, має пробіл у назві, вам потрібно буде використати дужкову нотацію.

Однак дужкову нотацію можна використовувати і для властивостей об’єкта без пробілів.

Ось приклад використання дужкової нотації для читання властивості об’єкта:

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]` буде рядком `Kirk`, `myObj['More Space']` буде рядком `Spock`, а `myObj["NoSpace"]` буде рядком `USS Enterprise`.

Зверніть увагу, що назви властивостей із пробілами повинні бути в лапках (одинарних або подвійних).

# --instructions--

Прочитайте значення властивостей `an entree` і `the drink` об’єкта `testObj` за допомогою дужкової нотації та призначте їх відповідно `entreeValue` і `drinkValue`.

# --hints--

`entreeValue` має бути рядком

```js
assert(typeof entreeValue === 'string');
```

Значення `entreeValue` має бути рядком `hamburger`

```js
assert(entreeValue === 'hamburger');
```

`drinkValue` має бути рядком

```js
assert(typeof drinkValue === 'string');
```

Значення `drinkValue` має бути рядком `water`

```js
assert(drinkValue === 'water');
```

Ви повинні використати дужкову нотацію двічі

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
