---
id: 56533eb9ac21ba0edf2244cd
title: Доступ до вкладених масивів
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

Як ми бачили в попередніх прикладах, об’єкти можуть містити як вкладені об’єкти, так і вкладені масиви. Подібно до доступу до вкладених об’єктів, дужкову нотацію масиву можна ланцюжково використовувати для доступу до вкладених масивів.

Ось приклад того, як отримати доступ до вкладеного масиву:

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` буде рядком `Fluffy`, а `ourPets[1].names[0]` буде рядком `Spot`.

# --instructions--

Використовуючи крапкову та дужкову нотації, встановіть змінну `secondTree` у другий елемент масиву `list` з другого об’єкта масиву `myPlants`.

# --hints--

`secondTree` має дорівнювати рядку `pine`.

```js
assert(secondTree === 'pine');
```

Ваш код має використовувати крапкову та дужкову нотації для доступу до `myPlants`.

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
