---
id: 56533eb9ac21ba0edf2244b7
title: Конкатенація рядків за допомогою оператора плюс
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

У JavaScript, коли оператор `+` використовується з рядковим значенням `String`, його називають оператором <dfn>конкатенації</dfn>. Ви можете створити новий рядок, <dfn>конкатенуючи</dfn> інші рядки разом.

**Приклад**

```js
'My name is Alan,' + ' I concatenate.'
```

**Примітка:** Зверніть увагу на пробіли. Конкатенація не додає пробіли між рядками, тому вам потрібно додати їх самостійно.

Приклад:

```js
const ourStr = "I come first. " + "I come second.";
```

Рядок `I come first. I come second.` буде виведений у консоль.
# --instructions--

Створіть `myStr` з рядків `This is the start.` та `This is the end.` за допомогою оператора `+`. Обов’язково додайте пробіл між двома рядками.

# --hints--

`myStr` повинен містити один символ пробілу між двома рядками.

```js
assert(/start\. This/.test(myStr));
```

`myStr` повинен мати значення рядка `This is the start. This is the end.`

```js
assert(myStr === 'This is the start. This is the end.');
```

Ви повинні використати оператор `+` для створення `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr` слід створити за допомогою ключового слова `const`.

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

Ви повинні призначити результат змінній `myStr`.

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
