---
id: 56533eb9ac21ba0edf2244a8
title: Збереження значень за допомогою оператора присвоєння
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

У JavaScript ви можете зберегти значення у змінній за допомогою <dfn>оператора присвоєння</dfn> (`=`).

```js
myVariable = 5;
```

Це призначає `Number` значення `5` змінній `myVariable`.

Якщо справа від оператора `=` є якісь обчислення, вони виконуються перед тим, як значення буде призначене змінній зліва від оператора.

```js
var myVar;
myVar = 5;
```

Спочатку цей код створює змінну з іменем `myVar`. Потім код призначає `5` змінній `myVar`. Тепер, якщо `myVar` з’явиться знову в коді, програма буде трактувати її так, ніби це `5`.

# --instructions--

Призначте значення `7` змінній `a`.

# --hints--

Не змінюйте код вище вказаного коментаря.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a` має мати значення 7.

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
