---
id: 56533eb9ac21ba0edf2244ed
title: Додавання змінних до рядків
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

Так само, як ми можемо створювати рядок на кількох лініях із рядкових <dfn>літералів</dfn>, ми також можемо додавати змінні до рядка за допомогою оператора плюс дорівнює (`+=`).

Приклад:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` матиме значення `freeCodeCamp is awesome!`.

# --instructions--

Встановіть `someAdjective` у рядок довжиною щонайменше 3 символи та додайте його до `myStr` за допомогою оператора `+=`.

# --hints--

`someAdjective` має бути рядком довжиною щонайменше 3 символи.

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

Ви повинні додати `someAdjective` до `myStr` за допомогою оператора `+=`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
