---
id: 56533eb9ac21ba0edf2244b9
title: Створення рядків за допомогою змінних
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

Іноді потрібно створити рядок. Використовуючи оператор конкатенації (`+`), ви можете вставити одну або кілька змінних у рядок, який створюєте.

приклад:

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` матиме значення рядка `Hello, our name is freeCodeCamp, how are you?`.

# --instructions--

Встановіть `myName` у рядок, що дорівнює вашому імені, і створіть `myStr`, розмістивши `myName` між рядками `My name is ` і ` and I am well!`.

# --hints--

`myName` має бути рядком довжиною щонайменше 3 символи.

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

Використайте два оператори `+`, щоб створити `myStr` з `myName` всередині.

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
