---
id: 56533eb9ac21ba0edf2244b8
title: Конкатенація рядків за допомогою оператора плюс рівно
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

Ми також можемо використовувати оператор `+=` для <dfn>конкатенації</dfn> рядка в кінець існуючої змінної рядка. Це дуже корисно, щоб розбити довгий рядок на кілька ліній.

**Примітка:** Слідкуйте за пробілами. Конкатенація не додає пробіли між об’єднаними рядками, тому вам потрібно додавати їх самостійно.

Приклад:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` тепер має значення рядка `I come first. I come second.`.

# --instructions--

Побудуйте `myStr` на кількох лініях, конкатенуючи ці два рядки: `This is the first sentence.` і `This is the second sentence.` за допомогою оператора `+=`. Використовуйте оператор `+=` так само, як показано в прикладі, і обов’язково додайте пробіл між двома рядками. Почніть із призначення першого рядка `myStr`, а потім додайте другий рядок.

# --hints--

`myStr` має містити один символ пробілу між двома рядками.

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` має мати значення рядка `This is the first sentence. This is the second sentence.`

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

Ви маєте використовувати оператор `+=` для побудови `myStr`.

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
