---
id: 56533eb9ac21ba0edf2244a9
title: Ініціалізація змінних оператором присвоєння
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

Загальною практикою є <dfn>ініціалізація</dfn> змінної початковим значенням у тій же лінії, де її оголошують.

```js
var myVar = 0;
```

Створює нову змінну з назвою `myVar` і призначає їй ініціальне значення `0`.

# --instructions--

Оголосіть змінну `a` за допомогою `var` та ініціалізуйте її значенням `9`.

# --hints--

Ви повинні ініціалізувати `a` значенням `9`.

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
