---
id: 56533eb9ac21ba0edf2244ad
title: Зменшення числа на 1 за допомогою JavaScript
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

Ви легко можете <dfn>зменшити</dfn> або знизити значення змінної на одиницю за допомогою оператора `--`.

```js
i--;
```

еквівалентний

```js
i = i - 1;
```

**Примітка:** Весь рядок стає `i--;`, тому знак «дорівнює» не потрібен.

# --instructions--

Змініть код, щоб використати оператор `--` для `myVar`.

# --hints--

`myVar` має дорівнювати `10`.

```js
assert(myVar === 10);
```

Потрібно змінити `myVar = myVar - 1;`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

Не слід призначати `myVar` значення `10`.

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

Використайте оператор `--` для `myVar`.

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

Не змінюйте код вище вказаного коментаря.

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
