---
id: 56533eb9ac21ba0edf2244d1
title: Порівняння з оператором строгої рівності
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

Оператор строгої рівності (`===`) є аналогом оператора рівності (`==`). Однак, на відміну від оператора рівності, який намагається привести обидва порівнювані значення до загального типу, оператор строгої рівності не виконує перетворення типів.

Якщо порівнювані значення мають різні типи, вони вважаються нерівними, і оператор строгої рівності поверне хиба.

**Приклади**

```js
3 ===  3  // true
3 === '3' // false
```

У другому прикладі `3` має тип `Number`, а `'3'` — тип `String`.

# --instructions--

Використайте оператор строгої рівності в інструкції `if`, щоб функція повертала рядок `Equal`, коли `val` строго дорівнює `7`.

# --hints--

`testStrict(10)` має повертати рядок `Not Equal`

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)` має повертати рядок `Equal`

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")` має повертати рядок `Not Equal`

```js
assert(testStrict('7') === 'Not Equal');
```

Ви маєте використати оператор `===`

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
