---
id: 56533eb9ac21ba0edf2244ae
title: Знаходження остачі в JavaScript
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

Оператор <dfn>остача</dfn> `%` дає остачу від ділення двох чисел.

**Приклад**

<pre>
5 % 2 = 1
5 / 2 = 2 остача 1
2 * 2 = 4
5 - 4 = 1
</pre>

**Використання**  
У математиці число можна перевірити на парність або непарність, перевіряючи остачу від ділення числа на `2`. Парні числа мають остачу `0`, а непарні — остачу `1`.

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**Примітка:** Оператор <dfn>остача</dfn> іноді помилково називають оператором модуля. Він дуже схожий на модуль, але не працює правильно з від’ємними числами.

# --instructions--

Присвойте `remainder` остачу від ділення `11` на `3` за допомогою оператора <dfn>остача</dfn> (`%`).

# --hints--

Змінна `remainder` має бути ініціалізована

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

Значення `remainder` має бути `2`

```js
assert(remainder === 2);
```

Ви маєте використати оператор `%`

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
