---
id: 56533eb9ac21ba0edf2244bb
title: Заповнення пропусків у реченнях
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

Вам надають речення з пропущеними словами, такими як іменники, дієслова, прикметники та прислівники. Ви повинні заповнити пропуски словами на свій вибір так, щоб утворене речення мало сенс.

Розглянемо це речення:

```md
It was really ____, and we ____ ourselves ____.
```

У цьому реченні є три пропуски — прикметник, дієслово та прислівник, і ми можемо додати слова на свій вибір, щоб завершити його. Потім ми можемо призначити готове речення змінній так:

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

У цьому завданні ми надаємо вам іменник, дієслово, прикметник і прислівник. Вам потрібно скласти повне речення, використовуючи слова на свій вибір разом із наданими словами.

Вам потрібно буде використати оператор конкатенації рядків `+` для створення нового рядка, використовуючи надані змінні: `myNoun`, `myAdjective`, `myVerb` та `myAdverb`. Потім ви призначите утворений рядок змінній `wordBlanks`. Ви не повинні змінювати слова, призначені цим змінним.

Також потрібно врахувати пробіли у вашому рядку, щоб фінальне речення містило пробіли між усіма словами. Результат має бути повним реченням.

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` має бути рядком.

```js
assert(typeof wordBlanks === 'string');
```

Ви не повинні змінювати значення, призначені `myNoun`, `myVerb`, `myAdjective` або `myAdverb`.

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

Ви не повинні безпосередньо використовувати значення `dog`, `ran`, `big` або `quickly` для створення `wordBlanks`.

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` має містити всі слова, призначені змінним `myNoun`, `myVerb`, `myAdjective` і `myAdverb`, розділені несловесними символами (а також будь-якими додатковими словами на ваш вибір).

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
