---
id: 56533eb9ac21ba0edf2244c0
title: Глобальна та локальна область видимості у функціях
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

Можна мати як <dfn>локальні</dfn>, так і <dfn>глобальні</dfn> змінні з однаковим ідентифікатором. Коли це трапляється, локальна змінна має пріоритет над глобальною.

У цьому прикладі:

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

Функція `myFun` поверне рядок `Head`, тому що присутня локальна версія змінної.

# --instructions--

Додайте локальну змінну у функцію `myOutfit`, щоб замістити значення `outerWear` рядком `sweater`.

# --hints--

Ви не повинні змінювати значення глобальної `outerWear`.

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit` має повертати рядок `sweater`.

```js
assert(myOutfit() === 'sweater');
```

Ви не повинні змінювати оператор return.

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
