---
id: 56533eb9ac21ba0edf2244c7
title: Доступ до властивостей об’єкта за допомогою точкової нотації
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

Існує два способи доступу до властивостей об’єкта: точкова нотація (`.`) і дужкова нотація (`[]`), подібно до масиву.

Точкову нотацію ви використовуєте, коли заздалегідь знаєте назву властивості, до якої хочете отримати доступ.

Ось приклад використання точкової нотації (`.`) для читання властивості об’єкта:

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` матиме значення рядка `val1`, а `prop2val` матиме значення рядка `val2`.

# --instructions--

Прочитайте значення властивостей `testObj` за допомогою точкової нотації. Присвойте змінній `hatValue` властивість `hat` об’єкта, а змінній `shirtValue` — властивість `shirt` об’єкта.

# --hints--

`hatValue` має бути рядком

```js
assert(typeof hatValue === 'string');
```

Значення `hatValue` має бути рядком `ballcap`

```js
assert(hatValue === 'ballcap');
```

`shirtValue` має бути рядком

```js
assert(typeof shirtValue === 'string');
```

Значення `shirtValue` має бути рядком `jersey`

```js
assert(shirtValue === 'jersey');
```

Ви маєте використати точкову нотацію двічі

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
