---
id: 56533eb9ac21ba0edf2244c9
title: Доступ до властивостей об’єкта за допомогою змінних
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

Ще одним способом використання дужкової нотації для об’єктів є доступ до властивості, яка зберігається як значення змінної. Це дуже корисно для перебору властивостей об’єкта або при доступі до таблиці пошуку.

Ось приклад використання змінної для доступу до властивості:

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

У консолі буде виведено рядок `Doberman`.

Зверніть увагу, що ми *не* використовуємо лапки навколо імені змінної, коли звертаємося до властивості, бо ми використовуємо *значення* змінної, а не її *ім’я*.

# --instructions--

Призначте змінній `playerNumber` значення `16`. Потім використайте цю змінну, щоб знайти ім’я гравця і призначити його `player`.

# --hints--

`playerNumber` має бути числом

```js
assert(typeof playerNumber === 'number');
```

Змінна `player` має бути рядком

```js
assert(typeof player === 'string');
```

Значення `player` має бути рядком `Montana`

```js
assert(player === 'Montana');
```

Використайте дужкову нотацію для доступу до `testObj`

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

Не призначайте значення `Montana` змінній `player` безпосередньо.

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

У дужковій нотації має використовуватися змінна `playerNumber`

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
