---
id: 56533eb9ac21ba0edf2244bf
title: Локальна область видимості та функції
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

Змінні, які оголошені всередині функції, а також параметри функції, мають <dfn>локальну</dfn> область видимості. Це означає, що вони видимі лише всередині цієї функції.

Ось функція `myTest` з локальною змінною під назвою `loc`.

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

Виклик функції `myTest()` виведе рядок `foo` у консоль. Лінія `console.log(loc)` (поза межами функції `myTest`) викличе помилку, оскільки `loc` не визначена поза межами функції.

# --instructions--

Редактор має два `console.log` щоб допомогти вам бачити, що відбувається. Перевіряйте консоль під час написання коду, щоб бачити, як вона змінюється. Оголосіть локальну змінну `myVar` всередині `myLocalScope` і запустіть тести.

**Примітка:** Консоль все одно відобразить `ReferenceError: myVar is not defined`, але це не спричинить помилку в тестах.

# --hints--

Код не повинен містити глобальну змінну `myVar`.

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

Ви повинні додати локальну змінну `myVar`.

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
