---
id: 56533eb9ac21ba0edf2244af
title: Складене присвоєння зі збільшенням
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

У програмуванні часто використовують присвоєння, щоб змінити вміст змінної. Пам’ятайте, що все праворуч від знака рівності спочатку обчислюється, тому ми можемо сказати:

```js
myVar = myVar + 5;
```

щоб додати `5` до `myVar`. Оскільки це дуже поширений шаблон, існують оператори, які виконують і математичну операцію, і присвоєння за один крок.

Одним із таких операторів є оператор `+=`.

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

У консолі буде виведено `6`.

# --instructions--

Перетворіть присвоєння для `a`, `b` і `c` так, щоб вони використовували оператор `+=`.

# --hints--

`a` має дорівнювати `15`.

```js
assert(a === 15);
```

`b` має дорівнювати `26`.

```js
assert(b === 26);
```

`c` має дорівнювати `19`.

```js
assert(c === 19);
```

Ви маєте використати оператор `+=` для кожної змінної.

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

Не змінюйте код вище за вказаний коментар.

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
