---
id: 56533eb9ac21ba0edf2244ac
title: Збільшення числа на 1 у JavaScript
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

Ви легко можете <dfn>збільшити</dfn> або додати одиницю до змінної за допомогою оператора `++`.

```js
i++;
```

еквівалентний

```js
i = i + 1;
```

**Примітка:** Весь рядок стає `i++;`, тому знак «дорівнює» не потрібен.

# --instructions--

Змініть код, щоб використати оператор `++` для `myVar`.

# --hints--

`myVar` має дорівнювати `88`.

```js
assert(myVar === 88);
```

Вам не слід використовувати оператор присвоєння.

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

Вам слід використати оператор `++`.

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

Не змінюйте код вище вказаного коментаря.

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
