---
id: 56533eb9ac21ba0edf2244be
title: Глобальна область видимості та функції
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

В JavaScript <dfn>область видимості</dfn> означає видимість змінних. Змінні, які визначені поза блоком функції, мають <dfn>глобальну</dfn> область видимості. Це означає, що їх можна бачити у будь-якій частині вашого JavaScript коду.

Змінні, які оголошені без ключових слів `let` або `const`, автоматично створюються в області видимості `global`. Це може спричинити небажані наслідки в інших частинах вашого коду або при повторному запуску функції. Ви завжди повинні оголошувати свої змінні за допомогою ключових слів `let` або `const`.

# --instructions--

Використовуючи `let` або `const`, оголосіть глобальну змінну з іменем `myGlobal` поза межами будь-якої функції. Ініціалізуйте її значенням `10`.

Всередині функції `fun1` призначте `5` змінній `oopsGlobal` ***без*** використання ключових слів `var`, `let` або `const`.

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal` має бути визначена

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal` має мати значення `10`

```js
assert(myGlobal === 10);
```

`myGlobal` має бути оголошена з використанням ключових слів `let` або `const`

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal` має бути глобальною змінною і мати значення `5`

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
