---
id: 56533eb9ac21ba0edf2244d0
title: Порівняння з оператором рівності
challengeType: 1
forumTopicId: 16784
dashedName: comparison-with-the-equality-operator
---

# --description--

У JavaScript існує багато <dfn>операторів порівняння</dfn>. Всі ці оператори повертають булеве `true` або `false` значення.

Найпростіший оператор — це оператор рівності `==`. Оператор рівності порівнює два значення і повертає `true`, якщо вони еквівалентні, або `false`, якщо ні. Зверніть увагу, що рівність відрізняється від присвоєння (`=`), яке призначає значення справа від оператора змінній зліва.

```js
function equalityTest(myVal) {
  if (myVal == 10) {
    return "Equal";
  }
  return "Not Equal";
}
```

Якщо `myVal` дорівнює `10`, оператор рівності поверне `true`, тому код у фігурних дужках виконається, а функція поверне `Equal`. Інакше функція поверне `Not Equal`. Щоб JavaScript міг порівняти два різні <dfn>типи даних</dfn> (наприклад, `numbers` і `strings`), він повинен перетворити один тип на інший. Це називається перетворенням типів. Після цього він може порівнювати значення так:

```js
1   ==  1  // true
1   ==  2  // false
1   == '1' // true
"3" ==  3  // true
```

# --instructions--

Додайте оператор рівності на вказану лінію, щоб функція повертала рядок `Equal`, коли `val` еквівалентний `12`.

# --hints--

`testEqual(10)` має повертати рядок `Not Equal`

```js
assert(testEqual(10) === 'Not Equal');
```

`testEqual(12)` має повертати рядок `Equal`

```js
assert(testEqual(12) === 'Equal');
```

`testEqual("12")` має повертати рядок `Equal`

```js
assert(testEqual('12') === 'Equal');
```

Ви маєте використати оператор `==`

```js
assert(__helpers.removeJSComments(code).match(/==/g) && !__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function testEqual(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testEqual(10);
```

# --solutions--

```js
function testEqual(val) {
  if (val == 12) {
    return "Equal";
  }
  return "Not Equal";
}
```
