---
id: 56533eb9ac21ba0edf2244b6
title: Керовані послідовності в рядках
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

Кавички — це не єдині символи, які можна <dfn>екранувати</dfn> всередині рядка. Керовані послідовності дозволяють вам використовувати символи, які інакше не можна було б використати в рядку.

<table><thead><tr><th>Код</th><th>Вихідні дані</th></tr></thead><tbody><tr><td><code>\'</code></td><td>одинарна кавичка</td></tr><tr><td><code>\"</code></td><td>подвійна кавичка</td></tr><tr><td><code>\\</code></td><td>зворотна скісна риска</td></tr><tr><td><code>\n</code></td><td>нова лінія</td></tr><tr><td><code>\t</code></td><td>табуляція</td></tr><tr><td><code>\r</code></td><td>повернення каретки</td></tr><tr><td><code>\b</code></td><td>зворотний пробіл</td></tr><tr><td><code>\f</code></td><td>переведення форми</td></tr></tbody></table>

*Зверніть увагу, що сама зворотна скісна риска має бути екранована, щоб відобразитися як зворотна скісна риска.*

# --instructions--

Призначте наступні три рядки тексту в одну змінну `myStr`, використовуючи керовані послідовності.

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

Вам потрібно буде використати керовані послідовності, щоб правильно вставити спеціальні символи. Також потрібно дотримуватися пробілів так, як показано вище, без пробілів між керованими послідовностями або словами.

**Примітка:** Відступ для `SecondLine` досягається за допомогою керованого символу табуляції, а не пробілів.

# --hints--

`myStr` не повинен містити пробілів

```js
assert(!/ /.test(myStr));
```

`myStr` повинен містити рядки `FirstLine`, `SecondLine` і `ThirdLine` (пам’ятайте про чутливість до регістру)

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine` має бути після символу нової лінії `\n`

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr` повинен містити символ табуляції `\t`, який йде після символу нової лінії

```js
assert(/\n\t/.test(myStr));
```

`SecondLine` має передувати символ зворотної скісної риски `\`

```js
assert(/\\SecondLine/.test(myStr));
```

Між `SecondLine` і `ThirdLine` має бути символ нової лінії

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr` повинен містити лише символи, показані в інструкціях

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
