---
id: 56533eb9ac21ba0edf2244bc
title: список покупок
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

Створіть список покупок у змінній `myList`. Список має бути багатовимірним масивом, що містить кілька підмасивів.

Перший елемент у кожному підмасиві повинен містити рядок із назвою елемента. Другий елемент має бути числом, що представляє кількість, наприклад

```js
["Chocolate Bar", 15]
```

У списку має бути щонайменше 5 підмасивів.

# --hints--

`myList` має бути масивом.

```js
assert(Array.isArray(myList));
```

Перші елементи у кожному з ваших підмасивів мають бути рядками.

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

Другі елементи у кожному з ваших підмасивів мають бути числами.

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

У вашому списку має бути щонайменше 5 елементів.

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
