---
id: 56533eb9ac21ba0edf2244de
title: додавання опції за замовчуванням у switch-інструкціях
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

У `switch` інструкції ви можете не мати змоги вказати всі можливі значення як `case` інструкції. Замість цього ви можете додати `default` інструкцію, яка буде виконана, якщо не знайдеться відповідних `case` інструкцій. Уявіть це як фінальну `else` інструкцію в ланцюжку `if/else`.

`default` інструкція має бути останньою.

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

Напишіть switch-інструкцію, щоб встановити `answer` для таких умов:  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` має повертати рядок `apple`

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` має повертати рядок `bird`

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` має повертати рядок `cat`

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` має повертати рядок `stuff`

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` має повертати рядок `stuff`

```js
assert(switchOfStuff(4) === 'stuff');
```

Ви не повинні використовувати жодних `if` або `else` інструкцій

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Ви повинні використати `default` інструкцію

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

У вас має бути принаймні 3 `break` інструкції

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
