---
id: 56533eb9ac21ba0edf2244ba
title: Розуміння незмінності рядків
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

У JavaScript значення `String` є <dfn>незмінними</dfn>, що означає, що їх не можна змінити після створення.

Наприклад, наступний код викличе помилку, тому що букву `B` у рядку `Bob` не можна змінити на букву `J`:

```js
let myStr = "Bob";
myStr[0] = "J";
```

Зверніть увагу, що це *не* означає, що `myStr` не можна перепризначити. Єдиний спосіб змінити `myStr` — це призначити їй нове значення, як у цьому прикладі:

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

Виправте присвоєння `myStr`, щоб воно містило рядкове значення `Hello World`, використовуючи підхід, показаний у прикладі вище.

# --hints--

`myStr` має мати значення рядка `Hello World`.

```js
assert(myStr === 'Hello World');
```

Не змінюйте код вище вказаного коментаря.

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
