---
id: 56533eb9ac21ba0edf2244c3
title: Присвоєння з поверненим значенням
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

Якщо ви пам’ятаєте з нашої розмови про <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">збереження значень за допомогою оператора присвоєння</a>, усе, що знаходиться праворуч від знака «дорівнює», обчислюється перед тим, як значення буде призначене. Це означає, що ми можемо взяти повернене значення функції і призначити його змінній.

Припустімо, ми визначили функцію `sum`, яка додає два числа.

```js
ourSum = sum(5, 12);
```

Виклик функції `sum` з аргументами `5` і `12` дає повернене значення `17`. Це повернене значення призначається змінній `ourSum`.

# --instructions--

Викличте функцію `processArg` з аргументом `7` і призначте її повернене значення змінній `processed`.

# --hints--

`processed` має мати значення `2`

```js
assert(processed === 2);
```

Ви повинні призначити `processArg` змінній `processed`

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
