---
id: 56533eb9ac21ba0edf2244db
title: Вступ до операторів else if
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

Якщо потрібно опрацювати кілька умов, ви можете поєднувати оператори `if` з операторами `else if`.

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

Перетворіть логіку так, щоб використовувати оператори `else if`.

# --hints--

У вас має бути принаймні два оператори `else`

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

У вас має бути принаймні два оператори `if`

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

Для кожного блоку коду `if else` мають бути відкриваючі та закриваючі фігурні дужки.

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

`testElseIf(0)` має повертати рядок `Smaller than 5`

```js
assert(testElseIf(0) === 'Smaller than 5');
```

`testElseIf(5)` має повертати рядок `Between 5 and 10`

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

`testElseIf(7)` має повертати рядок `Between 5 and 10`

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

`testElseIf(10)` має повертати рядок `Between 5 and 10`

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

`testElseIf(12)` має повертати рядок `Greater than 10`

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
