---
id: 56533eb9ac21ba0edf2244ab
title: Розуміння регістру в назвах змінних
challengeType: 1
forumTopicId: 18334
dashedName: understanding-case-sensitivity-in-variables
---

# --description--

У JavaScript усі змінні та імена функцій чутливі до регістру. Це означає, що велика чи мала літера має значення.

`MYVAR` не те саме, що `MyVar` або `myvar`. Можна мати кілька різних змінних з однаковим ім’ям, але різним регістром. Рішуче рекомендується для ясності *не використовувати* цю особливість мови.

**Найкраща практика**

Пишіть імена змінних у JavaScript у <dfn>верблюжому регістрі</dfn>. У <dfn>верблюжому регістрі</dfn> багатослівні імена змінних починаються з нижнього регістру, а перша літера кожного наступного слова — з великої.

**Приклади:**

```js
var someVariable;
var anotherVariableName;
var thisVariableNameIsSoLong;
```

# --instructions--

Змініть існуючі оголошення та присвоєння так, щоб їхні імена були у <dfn>верблюжому регістрі</dfn>.

Не створюйте нових змінних.

# --hints--

`studlyCapVar` має бути оголошеною і мати значення `10`.

```js
assert(typeof studlyCapVar !== 'undefined' && studlyCapVar === 10);
```

`properCamelCase` має бути оголошеною і мати значення рядка `A String`.

```js
assert(
  typeof properCamelCase !== 'undefined' && properCamelCase === 'A String'
);
```

`titleCaseOver` має бути оголошеною і мати значення `9000`.

```js
assert(typeof titleCaseOver !== 'undefined' && titleCaseOver === 9000);
```

`studlyCapVar` має використовувати верблюжий регістр і в оголошенні, і в присвоєнні.

```js
assert(__helpers.removeJSComments(code).match(/studlyCapVar/g).length === 2);
```

`properCamelCase` має використовувати верблюжий регістр і в оголошенні, і в присвоєнні.

```js
assert(__helpers.removeJSComments(code).match(/properCamelCase/g).length === 2);
```

`titleCaseOver` має використовувати верблюжий регістр і в оголошенні, і в присвоєнні.

```js
assert(__helpers.removeJSComments(code).match(/titleCaseOver/g).length === 2);
```

# --seed--

## --seed-contents--

```js
// Variable declarations
var StUdLyCapVaR;
var properCamelCase;
var TitleCaseOver;

// Variable assignments
STUDLYCAPVAR = 10;
PRoperCAmelCAse = "A String";
tITLEcASEoVER = 9000;
```

# --solutions--

```js
var studlyCapVar;
var properCamelCase;
var titleCaseOver;

studlyCapVar = 10;
properCamelCase = "A String";
titleCaseOver = 9000;
```
