---
id: 56533eb9ac21ba0edf2244aa
title: Розуміння неініціалізованих змінних
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

Коли оголошують змінні в JavaScript, вони мають ініціальне значення `undefined`. Якщо виконати математичну операцію над `undefined` змінною, результат буде `NaN`, що означає <dfn>"Не число"</dfn>. Якщо конкатенувати рядок із `undefined` змінною, отримаєте <dfn>рядок</dfn> `undefined`.

# --instructions--

Ініціалізуйте три змінні `a`, `b` і `c` значеннями `5`, `10` і `"I am a"` відповідно, щоб вони не були `undefined`.

# --hints--

`a` має бути оголошена і мати фінальне значення `6`.

```js
assert(typeof a === 'number' && a === 6);
```

`b` має бути оголошена і мати фінальне значення `15`.

```js
assert(typeof b === 'number' && b === 15);
```

`c` не повинна містити `undefined` і має мати фінальне значення рядка `I am a String!`

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

Не змінюйте код нижче вказаного коментаря.

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
