---
id: 56533eb9ac21ba0edf2244dd
title: Вибір серед багатьох варіантів за допомогою інструкції switch
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

Якщо потрібно порівняти одне значення з багатьма варіантами, можна використати інструкцію <dfn>switch</dfn>. Інструкція `switch` порівнює значення з інструкціями <dfn>case</dfn>, які визначають різні можливі варіанти. Будь-які дійсні JavaScript-інструкції можуть виконуватися всередині блоку <dfn>case</dfn> і виконуватимуться від першого співпавшого `case` значення до зустрічі `break`.

Ось приклад інструкції `switch`:

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

Значення `case` перевіряються на сувору рівність (`===`). `break` каже JavaScript припинити виконання інструкцій. Якщо `break` пропущено, буде виконана наступна інструкція.

# --instructions--

Напишіть інструкцію switch, яка перевіряє `val` і встановлює `answer` для таких умов:  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

`caseInSwitch(1)` має мати значення рядка `alpha`

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)` має мати значення рядка `beta`

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)` має мати значення рядка `gamma`

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)` має мати значення рядка `delta`

```js
assert(caseInSwitch(4) === 'delta');
```

Не слід використовувати інструкції `if` або `else`

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

Повинно бути щонайменше 3 інструкції `break`

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
