---
id: 56533eb9ac21ba0edf2244b0
title: Складене присвоєння зі збільшеним відніманням
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

Як і оператор `+=`, `-=` віднімає число від змінної.

```js
myVar = myVar - 5;
```

відніме `5` від `myVar`. Це можна переписати так:

```js
myVar -= 5;
```

# --instructions--

Перетворіть присвоєння для `a`, `b` і `c` так, щоб використовувався оператор `-=`.

# --hints--

`a` має дорівнювати `5`.

```js
assert(a === 5);
```

`b` має дорівнювати `-6`.

```js
assert(b === -6);
```

`c` має дорівнювати `2`.

```js
assert(c === 2);
```

Ви маєте використати оператор `-=` для кожної змінної.

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

Не змінюйте код вище вказаного коментаря.

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
