---
id: 56533eb9ac21ba0edf2244c6
title: Стати в чергу
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

У комп’ютерних науках <dfn>черга</dfn> — це абстрактна <dfn>структура даних</dfn>, у якій елементи зберігаються в певному порядку. Нові елементи можна додавати в кінець черги, а старі елементи видаляються з початку черги.

# --instructions--

Напишіть функцію `nextInLine`, яка приймає масив (`arr`) і число (`item`) як аргументи.

Додайте число в кінець масиву, потім видаліть перший елемент масиву.

Функція `nextInLine` має повернути елемент, який було видалено.

# --hints--

`nextInLine([], 5)` має повертати число.

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)` має повертати `1`

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)` має повертати `2`

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)` має повертати `5`

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
