---
id: 56533eb9ac21ba0edf2244cc
title: Доступ до вкладених об’єктів
challengeType: 1
forumTopicId: 16161
dashedName: accessing-nested-objects
---

# --description--

До вкладених властивостей об’єктів можна отримати доступ, послідовно використовуючи точкову або дужкову нотацію.

Ось вкладений об’єкт:

```js
const ourStorage = {
  "desk": {
    "drawer": "stapler"
  },
  "cabinet": {
    "top drawer": { 
      "folder1": "a file",
      "folder2": "secrets"
    },
    "bottom drawer": "soda"
  }
};

ourStorage.cabinet["top drawer"].folder2;
ourStorage.desk.drawer;
```

`ourStorage.cabinet["top drawer"].folder2` буде рядком `secrets`, а `ourStorage.desk.drawer` буде рядком `stapler`.

# --instructions--

Отримайте доступ до об’єкта `myStorage` і призначте вміст властивості `glove box` змінній `gloveBoxContents`. Використовуйте точкову нотацію для всіх властивостей, де це можливо, інакше застосовуйте дужкову нотацію.

# --hints--

`gloveBoxContents` має дорівнювати рядку `maps`.

```js
assert(gloveBoxContents === 'maps');
```

Ваш код має використовувати точкову нотацію, де це можливо, щоб отримати доступ до `myStorage`.

```js
assert.match(code, /myStorage\.car\.inside/);
```

`gloveBoxContents` все ще має бути оголошено за допомогою `const`.

```js
assert.match(code, /const\s+gloveBoxContents\s*=/);
```

Ви не повинні змінювати об’єкт `myStorage`.

```js
const expectedMyStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
assert.deepStrictEqual(myStorage, expectedMyStorage);
```

# --seed--

## --seed-contents--

```js
const myStorage = {
  "car": {
    "inside": {
      "glove box": "maps",
      "passenger seat": "crumbs"
     },
    "outside": {
      "trunk": "jack"
    }
  }
};

const gloveBoxContents = undefined;
```

# --solutions--

```js
const myStorage = {
  "car":{
    "inside":{
      "glove box":"maps",
      "passenger seat":"crumbs"
    },
    "outside":{
      "trunk":"jack"
    }
  }
};
const gloveBoxContents = myStorage.car.inside["glove box"];
```
