---
id: 56533eb9ac21ba0edf2244bd
title: Передача значень у функції за допомогою аргументів
challengeType: 1
forumTopicId: 18254
dashedName: passing-values-to-functions-with-arguments
---

# --description--

<dfn>Параметри</dfn> — це змінні, які виступають як заповнювачі для значень, що передаються у функцію під час її виклику. Коли функцію визначають, зазвичай вона визначається разом з одним або кількома параметрами. Фактичні значення, які передаються (або <dfn>передаються</dfn>) у функцію під час її виклику, називаються <dfn>аргументами</dfn>.

Ось функція з двома параметрами, `param1` та `param2`:

```js
function testFun(param1, param2) {
  console.log(param1, param2);
}
```

Потім ми можемо викликати `testFun` так: `testFun("Hello", "World");`. Ми передали два рядкові аргументи, `Hello` та `World`. Всередині функції `param1` дорівнюватиме рядку `Hello`, а `param2` дорівнюватиме рядку `World`. Зверніть увагу, що ви можете викликати `testFun` знову з іншими аргументами, і параметри приймуть значення нових аргументів.

# --instructions--

<ol><li>Створіть функцію з назвою <code>functionWithArgs</code>, яка приймає два аргументи і виводить їх суму у консоль розробника.</li><li>Викличте функцію з двома числами як аргументами.</li></ol>

# --before-each--

```js
var logOutput;
```

# --hints--

`functionWithArgs` має бути функцією.

```js
assert(typeof functionWithArgs === 'function');
```

`functionWithArgs(1,2)` має виводити `3`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(1, 2);
assert(logOutput == 3);
```

`functionWithArgs(7,9)` має виводити `16`.

```js
console.log = function (message) {
  logOutput = message
};
functionWithArgs(7, 9);
assert(logOutput == 16);
```

Ви маєте викликати `functionWithArgs` з двома числами після того, як визначите її.

```js
assert(
  /functionWithArgs\([-+]?\d*\.?\d*,[-+]?\d*\.?\d*\)/.test(
    __helpers.removeJSComments(code).replace(/\s/g, '')
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function functionWithArgs(a, b) {
  console.log(a + b);
}
functionWithArgs(10, 5);
```
