---
id: 56533eb9ac21ba0edf2244d2
title: Порівняння з оператором нерівності
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

Оператор нерівності (`!=`) є протилежним оператору рівності. Нерівність означає «не дорівнює». Оператор нерівності повертає `false`, коли оператор рівності повернув би `true`, і *навпаки*. Як і оператор рівності, оператор нерівності під час порівняння перетворює типи даних значень.

**Приклади**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

Додайте оператор нерівності `!=` в інструкцію `if`, щоб функція повертала рядок `Not Equal`, коли `val` не є еквівалентним `99`.

# --hints--

`testNotEqual(99)` має повертати рядок `Equal`

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` має повертати рядок `Equal`

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` має повертати рядок `Not Equal`

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` має повертати рядок `Not Equal`

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` має повертати рядок `Not Equal`

```js
assert(testNotEqual('bob') === 'Not Equal');
```

Ви маєте використати оператор `!=`

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
