---
id: 56592a60ddddeae28f7aa8e1
title: Доступ до багатовимірних масивів за індексами
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

Один зі способів уявити собі <dfn>багатовимірний</dfn> масив — це *масив масивів*. Коли ви використовуєте дужки для доступу до свого масиву, перша пара дужок посилається на записи у зовнішньому (першому рівні) масиві, а кожна додаткова пара дужок — на наступний рівень записів всередині.

**Приклад**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

У цьому прикладі `subarray` має значення `[[10, 11, 12], 13, 14]`,
`nestedSubarray` має значення `[10, 11, 12]`, а `element` має значення `11`.

**Примітка:** Між ім’ям масиву та квадратними дужками не повинно бути пробілів, як у `array [0][0]`, і навіть таке `array [0] [0]` не дозволяється. Хоча JavaScript може обробити це правильно, це може збити з пантелику інших програмістів, які читають ваш код.

# --instructions--

Використовуючи дужкову нотацію, виберіть елемент з `myArray` так, щоб `myData` дорівнювало `8`.

# --hints--

`myData` має дорівнювати `8`.

```js
assert(myData === 8);
```

Ви повинні використовувати дужкову нотацію, щоб прочитати правильне значення з `myArray`.

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
