---
id: 5664820f61c48e80c9fa476c
title: Код для гольфу
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

У грі в гольф кожна лунка має `par`, тобто середню кількість `strokes`, які гравець очікується зробити, щоб загнати м'яч у лунку і завершити гру. Залежно від того, наскільки ваші `par` відрізняються від `strokes`, існує різна назва.

Вашій функції будуть передані `par` і `strokes` аргументи. Поверніть правильний рядок згідно з цією таблицею, яка перераховує штрихи за пріоритетом; зверху (найвищий) донизу (найнижчий):

<table><thead><tr><th>Штрихи</th><th>Повернути</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par` і `strokes` завжди будуть числовими і додатними. Ми додали масив усіх назв для вашої зручності.

# --hints--

`golfScore(4, 1)` має повертати рядок `Hole-in-one!`

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)` має повертати рядок `Eagle`

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)` має повертати рядок `Eagle`

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)` має повертати рядок `Birdie`

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)` має повертати рядок `Par`

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)` має повертати рядок `Hole-in-one!`

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)` має повертати рядок `Par`

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)` має повертати рядок `Bogey`

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)` має повертати рядок `Double Bogey`

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)` має повертати рядок `Go Home!`

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)` має повертати рядок `Go Home!`

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
