---
id: 5675e877dbd60be8ad28edc6
title: Ітерувати по масиву за допомогою циклу for
challengeType: 1
forumTopicId: 18216
dashedName: iterate-through-an-array-with-a-for-loop
---

# --description--

Загальне завдання в JavaScript — ітерувати по вмісту масиву. Один зі способів зробити це — використати цикл `for`. Цей код виведе кожен елемент масиву `arr` у консоль:

```js
const arr = [10, 9, 8, 7, 6];

for (let i = 0; i < arr.length; i++) {
   console.log(arr[i]);
}
```

Пам’ятайте, що масиви мають індексацію, починаючи з нуля, тобто останній індекс масиву — це `length - 1`. Умова для цього циклу — `i < arr.length`, яка зупиняє цикл, коли `i` дорівнює `length`. У цьому випадку остання ітерація — це `i === 4`, тобто коли `i` стає рівним `arr.length - 1` і виводить `6` у консоль. Потім `i` збільшується до `5`, і цикл завершується, бо `i < arr.length` є `false`.

# --instructions--

Оголосіть і ініціалізуйте змінну `total` значенням `0`. Використайте цикл `for`, щоб додати значення кожного елемента масиву `myArr` до `total`.

# --hints--

`total` слід оголосити і ініціалізувати значенням 0.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s*?total\s*=\s*0.*?;?/));
```

`total` має дорівнювати 20.

```js
assert(total === 20);
```

Ви маєте використати цикл `for`, щоб ітерувати по `myArr`.

```js
assert(/for\s*\(/g.test(__helpers.removeJSComments(code)) && /myArr\s*\[/g.test(__helpers.removeJSComments(code)));
```

Не слід намагатися безпосередньо призначити значення 20 змінній `total`.

```js
assert(!__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/total[=+-]0*[1-9]+/gm));
```

# --seed--

## --seed-contents--

```js
// Setup
const myArr = [2, 3, 4, 5, 6];

// Only change code below this line

```

# --solutions--

```js
const myArr = [2, 3, 4, 5, 6];
let total = 0;

for (let i = 0; i < myArr.length; i++) {
  total += myArr[i];
}
```
