---
id: 5690307fddb111c6084545d7
title: Логічний порядок в інструкціях if else
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

Порядок має значення в інструкціях `if`, `else if`.

Функція виконується зверху вниз, тому слід уважно стежити, яка інструкція йде першою.

Візьмемо для прикладу ці дві функції.

Ось перша:

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

А друга просто змінює порядок інструкцій:

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

Хоча ці дві функції виглядають майже ідентичними, якщо передати число в обидві, ми отримаємо різні вихідні дані.

```js
foo(0)
bar(0)
```

`foo(0)` поверне рядок `Less than one`, а `bar(0)` поверне рядок `Less than two`.

# --instructions--

Змініть порядок логіки у функції так, щоб вона повертала правильні інструкції у всіх випадках.

# --hints--

`orderMyLogic(4)` має повертати рядок `Less than 5`

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` має повертати рядок `Less than 10`

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` має повертати рядок `Greater than or equal to 10`

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
