---
id: 56bbb991ad1ed5201cd392d1
title: оновлення властивостей об’єкта
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

Після створення об’єкта JavaScript ви можете оновлювати його властивості в будь-який момент так само, як оновлюєте будь-яку іншу змінну. Для оновлення можна використовувати як крапкову, так і дужкову нотацію.

Наприклад, розглянемо `ourDog`:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

Оскільки він особливо щасливий пес, давайте змінимо його ім’я на рядок `Happy Camper`. Ось як ми оновлюємо властивість name його об’єкта: `ourDog.name = "Happy Camper";` або `ourDog["name"] = "Happy Camper";` Тепер, коли ми оцінюємо `ourDog.name`, замість `Camper` отримаємо його нове ім’я, `Happy Camper`.

# --instructions--

Оновіть властивість name об’єкта `myDog`. Змініть її ім’я з `Coder` на `Happy Coder`. Можна використовувати як крапкову, так і дужкову нотацію.

# --hints--

Ви маєте оновити властивість `myDog` об’єкта `name`, присвоївши їй рядок `Happy Coder`.

```js
assert(/happy coder/gi.test(myDog.name));
```

Визначення `myDog` редагувати не потрібно.

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
