---
id: 56bbb991ad1ed5201cd392cc
title: Маніпулюйте масивами за допомогою методу pop
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

Ще один спосіб змінити дані в масиві — це використати функцію `.pop()`.

`.pop()` використовується, щоб вилучити значення з кінця масиву. Ми можемо зберегти це вилучене значення, призначивши його змінній. Іншими словами, `.pop()` видаляє останній елемент із масиву і повертає цей елемент.

З масиву можна вилучати будь-який тип елементів — числа, рядки, навіть вкладені масиви.

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

Перший `console.log` виведе значення `6`, а другий — значення `[1, 4]`.

# --instructions--

Використайте функцію `.pop()`, щоб видалити останній елемент із `myArray` і призначити вилучене значення новій змінній `removedFromMyArray`.

# --hints--

`myArray` має містити лише `[["John", 23]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

Ви маєте використати `pop()` на `myArray`.

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` має містити лише `["cat", 2]`.

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
