---
id: 56bbb991ad1ed5201cd392d2
title: додати нові властивості до об’єкта JavaScript
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

Ви можете додати нові властивості до існуючих об’єктів JavaScript так само, як і змінювати їх.

Ось як ми додамо властивість `bark` до `ourDog`:

```js
ourDog.bark = "bow-wow";
```

або

```js
ourDog["bark"] = "bow-wow";
```

Тепер, коли ми оцінюємо `ourDog.bark`, отримаємо його гавкіт, `bow-wow`.

Приклад:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

Додайте властивість `bark` до `myDog` і встановіть її значенням звук собаки, наприклад "woof". Ви можете використовувати як крапкову, так і дужкову нотацію.

# --hints--

Ви повинні додати властивість `bark` до `myDog`.

```js
assert(myDog.bark !== undefined);
```

Ви не повинні додавати `bark` під час ініціалізації `myDog`.

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
