---
id: 56bbb991ad1ed5201cd392cf
title: Написання багаторазового JavaScript за допомогою функцій
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

У JavaScript ми можемо розділити свій код на багаторазові частини, які називаються <dfn>функціями</dfn>.

Ось приклад функції:

```js
function functionName() {
  console.log("Hello World");
}
```

Ви можете викликати або <dfn>запустити</dfn> цю функцію, використовуючи її ім’я, за яким слідують круглі дужки, ось так: `functionName();` Кожного разу, коли функція викликається, вона виводитиме повідомлення `Hello World` у консоль розробника. Увесь код між фігурними дужками буде виконуватися щоразу, коли функція викликається.

# --instructions--

<ol>
  <li>
    Створіть функцію з назвою <code>reusableFunction</code>, яка виводить рядок <code>Hi World</code> у консоль розробника.
  </li>
  <li>
    Викличте цю функцію.
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction` має бути функцією.

```js
assert(typeof reusableFunction === 'function');
```

Якщо `reusableFunction` викликається, вона має виводити рядок `Hi World` у консоль.

```js
assert(testConsole());
```

Ви маєте викликати `reusableFunction` після того, як вона буде визначена.

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
