---
id: 56bbb991ad1ed5201cd392cb
title: Маніпулюйте масивами за допомогою методу push
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

Простий спосіб додати дані в кінець масиву — це використати метод `push()`.

Метод `push()` приймає один або кілька <dfn>аргументів</dfn> і додає їх у кінець масиву в тому порядку, в якому вони передані. Він повертає нову довжину масиву.

Приклади:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1` тепер має значення `[1, 2, 3, 4, 5]`, а `arr2` має значення `["Stimpson", "J", "cat", ["happy", "joy"]]`.

# --instructions--

Додайте `["dog", 3]` у кінець змінної `myArray`.

# --hints--

`myArray` тепер має дорівнювати `[["John", 23], ["cat", 2], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
