---
id: 56bbb991ad1ed5201cd392cd
title: Маніпулюйте масивами за допомогою методу shift
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` завжди видаляє останній елемент масиву. А що, якщо потрібно видалити перший?

Ось тут і допомагає `.shift()`. Він працює так само, як `.pop()`, але замість останнього видаляє перший елемент.

Приклад:

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` матиме значення рядка `Stimpson`, а `ourArray` матиме `["J", ["cat"]]`.

# --instructions--

Використайте функцію `.shift()`, щоб видалити перший елемент з `myArray` і призначити "зсунутий" елемент новій змінній `removedFromMyArray`.

# --hints--

`myArray` тепер має дорівнювати `[["dog", 3]]`.

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` має містити `["John", 23]`.

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
