---
id: 56bbb991ad1ed5201cd392ce
title: Маніпулюйте масивами за допомогою методу unshift
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

Ви можете не лише `shift` елементи з початку масиву, а й `unshift` елементи на початок масиву, тобто додати елементи перед масивом.

`.unshift()` працює точно так само, як `.push()`, але замість додавання елемента в кінець масиву, `unshift()` додає елемент на початок масиву.

Приклад:

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

Після `shift` `ourArray` матиме значення `["J", "cat"]`. Після `unshift` `ourArray` матиме значення `["Happy", "J", "cat"]`.

# --instructions--

Додайте `["Paul", 35]` на початок змінної `myArray` за допомогою `unshift()`.

# --hints--

`myArray` тепер має містити `[["Paul", 35], ["dog", 3]]`.

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
