---
id: 56bbb991ad1ed5201cd392ca
title: доступ до даних масиву за індексами
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

Ми можемо отримати доступ до даних всередині масивів, використовуючи <dfn>індекси</dfn>.

Індекси масивів записуються у тій самій дужковій нотації, що й рядки, але замість вказування символу вони вказують на елемент масиву. Як і рядки, масиви використовують <dfn>індексацію, починаючи з нуля</dfn>, тому перший елемент масиву має індекс `0`.

<br>

**Приклад**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

Змінна `console.log(array[0])` виводить `50`, а `data` має значення `60`.

# --instructions--

Створіть змінну з назвою `myData` і присвойте їй перше значення `myArray`, використовуючи дужкову нотацію.

# --hints--

Змінна `myData` має дорівнювати першому значенню `myArray`.

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

Дані у змінній `myArray` потрібно отримати за допомогою дужкової нотації.

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
