---
id: 56bbb991ad1ed5201cd392d0
title: Створення об’єктів JavaScript
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

Ви, можливо, вже чули термін `object` раніше.

Об’єкти схожі на `arrays`, але замість того, щоб використовувати індекси для доступу та зміни їхніх даних, ви отримуєте доступ до даних в об’єктах через те, що називають `properties`.

Об’єкти корисні для зберігання даних у структурованому вигляді і можуть представляти реальні об’єкти, наприклад, кота.

Ось приклад об’єкта кота:

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

У цьому прикладі всі властивості збережені як рядки, такі як `name`, `legs` і `tails`. Однак ви також можете використовувати числа як властивості. Ви навіть можете опустити лапки для однословних рядкових властивостей, як показано нижче:

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

Проте, якщо у вашому об’єкті є будь-які властивості, що не є рядками, JavaScript автоматично приведе їх до рядкового типу.

# --instructions--

Створіть об’єкт, який представляє собаку на ім’я `myDog`, і який містить властивості `name` (рядок), `legs`, `tails` та `friends`.

Ви можете призначити цим властивостям будь-які значення, які хочете, за умови, що `name` є рядком, `legs` і `tails` — числами, а `friends` — масивом.

# --hints--

`myDog` має містити властивість `name`, і вона має бути `string`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` має містити властивість `legs`, і вона має бути `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` має містити властивість `tails`, і вона має бути `number`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` має містити властивість `friends`, і вона має бути `array`.

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` має містити лише всі задані властивості.

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
