---
id: 56bbb991ad1ed5201cd392d3
title: Видалення властивостей з об’єкта JavaScript
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

Ми також можемо видаляти властивості з об’єктів таким чином:

```js
delete ourDog.bark;
```

Приклад:

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

Після останньої лінії, показаної вище, `ourDog` виглядає так:

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

Видаліть властивість `tails` з `myDog`. Ви можете використати як крапкову, так і дужкову нотацію.

# --hints--

Ви повинні видалити властивість `tails` з `myDog`.

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

Ви не повинні змінювати налаштування `myDog`.

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
