---
id: 587d7b7e367417b2b2512b24
title: Використання умовного (тернарного) оператора
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<dfn>Умовний оператор</dfn>, також званий <dfn>тернарним оператором</dfn>, можна використовувати як однорядковий if-else вираз.

Синтаксис такий: `a ? b : c`, де `a` — це умова, `b` — код, який виконується, коли умова повертає `true`, а `c` — код, який виконується, коли умова повертає `false`.

Наступна функція використовує інструкцію `if/else` для перевірки умови:

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

Це можна переписати, використовуючи умовний оператор:

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

Використайте умовний оператор у функції `checkEqual`, щоб перевірити, чи два числа рівні чи ні. Функція має повертати або рядок `Equal`, або рядок `Not Equal`.

# --hints--

`checkEqual` має використовувати умовний оператор

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)` має повертати рядок `Not Equal`

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)` має повертати рядок `Equal`

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)` має повертати рядок `Not Equal`

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
