---
id: 587d7b87367417b2b2512b3f
title: Дослідіть відмінності між ключовими словами var і let
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

Одна з найбільших проблем оголошення змінних за допомогою ключового слова `var` полягає в тому, що ви легко можете перезаписати оголошення змінної:

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

У наведеному вище коді змінна `camper` спочатку оголошена як `James`, а потім перезаписана як `David`. У консолі відображається рядок `David`.

У невеликому застосунку ви можете не зіткнутися з такою проблемою. Але коли ваш код стає більшим, ви можете випадково перезаписати змінну, яку не мали наміру змінювати. Оскільки така поведінка не викликає помилку, пошук і виправлення помилок стає складнішим.

Ключове слово `let` було введено в ES6, великому оновленні JavaScript, щоб вирішити цю потенційну проблему з ключовим словом `var`. Ви дізнаєтеся про інші функції ES6 у наступних завданнях.

Якщо замінити `var` на `let` у наведеному вище коді, це призведе до помилки:

```js
let camper = "James";
let camper = "David";
```

Помилку можна побачити у консолі браузера.

Отже, на відміну від `var`, коли ви використовуєте `let`, змінна з таким самим іменем може бути оголошена лише один раз.

# --instructions--

Оновіть код так, щоб він використовував лише ключове слово `let`.

# --hints--

`var` не повинно бути в коді.

```js
assert.notMatch(code, /var/g);
```

`catName` має бути рядком `Oliver`.

```js
assert.equal(catName, 'Oliver');
```

`catSound` має бути рядком `Meow!`

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
