---
id: 587d7b7e367417b2b2512b22
title: Використання функції parseInt з основою числення
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

Функція `parseInt()` аналізує рядок і повертає ціле число. Вона приймає другий аргумент — основу числення, яка вказує систему числення числа в рядку. Основа числення може бути цілим числом від 2 до 36.

Виклик функції виглядає так:

```js
parseInt(string, radix);
```

Ось приклад:

```js
const a = parseInt("11", 2);
```

Змінна radix вказує, що `11` записане у двійковій системі, або в основі 2. Цей приклад конвертує рядок `11` у ціле число `3`.

# --instructions--

Використайте `parseInt()` у функції `convertToInteger`, щоб вона конвертувала двійкове число у ціле число і повертала його.

# --hints--

`convertToInteger` має використовувати функцію `parseInt()`

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` має повертати число

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` має повертати 19

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` має повертати 57

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` має повертати `NaN`

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
