---
id: 587d7b7e367417b2b2512b21
title: Використання кількох умовних (тернарних) операторів
challengeType: 1
forumTopicId: 301179
dashedName: use-multiple-conditional-ternary-operators
---

# --description--

У попередньому завданні ви використали один умовний оператор. Ви також можете ланцюжити їх, щоб перевірити кілька умов.

Наступна функція використовує інструкції `if`, `else if` та `else` для перевірки кількох умов:

```js
function findGreaterOrEqual(a, b) {
  if (a === b) {
    return "a and b are equal";
  }
  else if (a > b) {
    return "a is greater";
  }
  else {
    return "b is greater";
  }
}
```

Вищенаведену функцію можна переписати, використовуючи кілька умовних операторів:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" 
    : (a > b) ? "a is greater" 
    : "b is greater";
}
```

Вважається найкращою практикою форматувати кілька умовних операторів так, щоб кожна умова була на окремому рядку, як показано вище. Використання кількох умовних операторів без належного відступу може ускладнити читання вашого коду. Наприклад:

```js
function findGreaterOrEqual(a, b) {
  return (a === b) ? "a and b are equal" : (a > b) ? "a is greater" : "b is greater";
}
```

# --instructions--

У функції `checkSign` використайте кілька умовних операторів — дотримуючись рекомендованого формату, який використовується в `findGreaterOrEqual` — щоб перевірити, чи є число додатним, від’ємним або нулем. Функція має повертати `positive`, `negative` або `zero`.

# --hints--

`checkSign` має використовувати кілька умовних операторів

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?\s*?\?\s*?.+?\s*?:\s*?.+?/gi.test(__helpers.removeJSComments(code)));
```

`checkSign(10)` має повертати рядок `positive`. Зверніть увагу, що важлива велика чи мала літера

```js
assert(checkSign(10) === 'positive');
```

`checkSign(-12)` має повертати рядок `negative`. Зверніть увагу, що важлива велика чи мала літера

```js
assert(checkSign(-12) === 'negative');
```

`checkSign(0)` має повертати рядок `zero`. Зверніть увагу, що важлива велика чи мала літера

```js
assert(checkSign(0) === 'zero');
```

# --seed--

## --seed-contents--

```js
function checkSign(num) {

}

checkSign(10);
```

# --solutions--

```js
function checkSign(num) {
  return (num > 0) ? 'positive' : (num < 0) ? 'negative' : 'zero';
}
```
