---
id: 598e8944f009e646fc236146
title: "Розуміння значення undefined, що повертається з функції"
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

Функція може містити інструкцію `return`, але це не обов’язково. Якщо функція не має інструкції `return`, коли ви її викликаєте, функція виконує внутрішній код, але повернене значення буде `undefined`.

**Приклад**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` — це функція без інструкції `return`. Функція змінить глобальну змінну `sum`, але повернене значення функції — `undefined`.

# --instructions--

Створіть функцію `addFive` без аргументів. Ця функція додає 5 до змінної `sum`, але її повернене значення — `undefined`.

# --hints--

`addFive` має бути функцією.

```js
assert(typeof addFive === 'function');
```

Після запуску обох функцій змінна `sum` має дорівнювати `8`.

```js
assert(sum === 8);
```

Повернене значення з `addFive` має бути `undefined`.

```js
assert(addFive() === undefined);
```

Всередині функції `addFive` потрібно додати `5` до змінної `sum`.

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
