---
id: 599a789b454f2bbd91a3ff4d
title: Практика порівняння різних значень
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

У двох останніх завданнях ми вивчили оператор рівності (`==`) та оператор строгої рівності (`===`). Давайте швидко повторимо і потренуємося ще раз використовувати ці оператори.

Якщо значення, які порівнюють, не одного типу, оператор рівності виконає перетворення типу, а потім оцінить значення. Однак оператор строгої рівності порівнює і тип даних, і значення без жодного перетворення.

**Приклади**

`3 == '3'` повертає `true`, тому що JavaScript виконує перетворення типу з рядка на число. `3 === '3'` повертає `false`, тому що типи різні і перетворення типу не виконується.

**Примітка:** У JavaScript ви можете визначити тип змінної або значення за допомогою оператора `typeof`, як показано нижче:

```js
typeof 3
typeof '3'
```

`typeof 3` повертає рядок `number`, а `typeof '3'` повертає рядок `string`.

# --instructions--

Функція `compareEquality` в редакторі порівнює два значення за допомогою оператора рівності. Змініть функцію так, щоб вона повертала рядок `Equal` лише тоді, коли значення строго рівні.

# --hints--

`compareEquality(10, "10")` має повертати рядок `Not Equal`

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` має повертати рядок `Not Equal`

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

Ви маєте використати оператор `===`

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
