---
id: 5cc0bd7a49b71cb96132e54c
title: Використовуйте рекурсію для створення діапазону чисел
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

Продовжуючи попереднє завдання, ми надаємо вам ще одну можливість створити рекурсивну функцію для розв’язання проблеми.

# --instructions--

Ми визначили функцію з іменем `rangeOfNumbers` з двома параметрами. Функція має повертати масив цілих чисел, який починається з числа, представленого параметром `startNum`, і закінчується числом, представленим параметром `endNum`. Початкове число завжди буде меншим або рівним кінцевому числу. Ваша функція має використовувати рекурсію, викликаючи сама себе, і не повинна використовувати цикли будь-якого виду. Вона також має працювати у випадках, коли `startNum` і `endNum` однакові.

# --hints--

Ваша функція має повертати масив.

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

Ваш код не повинен використовувати жодного синтаксису циклів (`for` або `while`) або функцій вищого порядку, таких як `forEach`, `map`, `filter` або `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers` має використовувати рекурсію (викликати сама себе) для розв’язання цього завдання.

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)` має повертати `[1, 2, 3, 4, 5]`.

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)` має повертати `[6, 7, 8, 9]`.

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)` має повертати `[4]`.

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

Глобальні змінні не повинні використовуватися для кешування масиву.

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
