---
id: 5cfa3679138e7d9595b9d9d4
title: Замінити цикли за допомогою рекурсії
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

Рекурсія — це концепція, за якою функція може бути виражена через саму себе. Щоб краще це зрозуміти, почніть із такого завдання: помножити перші `n` елементів масиву, щоб отримати добуток цих елементів. Використовуючи цикл `for`, ви могли б зробити це так:

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

Однак зверніть увагу, що `multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`. Це означає, що ви можете переписати `multiply` через саму себе і ніколи не використовувати цикл.

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

Рекурсивна версія `multiply` розбивається так. У <dfn>базовому випадку</dfn>, коли `n <= 0`, вона повертає 1. Для більших значень `n` вона викликає саму себе, але з `n - 1`. Цей виклик функції оцінюється так само, викликаючи `multiply` знову, доки не досягне `n <= 0`. У цей момент усі функції можуть повернути результат, і початкова `multiply` повертає відповідь.

**Примітка:** Рекурсивні функції повинні мати базовий випадок, коли вони повертають значення без повторного виклику функції (у цьому прикладі, коли `n <= 0`), інакше вони ніколи не завершать виконання.

# --instructions--

Напишіть рекурсивну функцію `sum(arr, n)`, яка повертає суму перших `n` елементів масиву `arr`.

# --hints--

`sum([1], 0)` має дорівнювати 0.

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)` має дорівнювати 2.

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)` має дорівнювати 9.

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

Ваш код не повинен покладатися на будь-які цикли (`for` або `while`) чи функції вищого порядку, такі як `forEach`, `map`, `filter` або `reduce`.

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

Використовуйте рекурсію для розв’язання цього завдання.

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
