---
id: 5ee127a03c3b35dd45426493
title: Присвоєння значення однієї змінної іншій
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

Після того, як значення призначене змінній за допомогою оператора <dfn>присвоєння</dfn>, ви можете призначити значення цієї змінної іншій змінній, використовуючи оператор <dfn>присвоєння</dfn>.

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

Вищенаведений код оголошує змінну `myVar` без значення, а потім призначає їй значення `5`. Далі оголошується змінна з ім’ям `myNum` без значення. Потім вміст `myVar` (який є `5`) призначається змінній `myNum`. Тепер `myNum` також має значення `5`.

# --instructions--

Призначте вміст `a` змінній `b`.

# --hints--

Не слід змінювати код вище вказаного коментаря.

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b` має мати значення `7`.

```js
assert(typeof b === 'number' && b === 7);
```

`a` слід призначити `b` за допомогою `=`.

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
