---
id: bd7123c9c441eddfaeb4bdef
title: Коментуйте свій JavaScript код
challengeType: 1
forumTopicId: 16783
dashedName: comment-your-javascript-code
---

# --description--

Коментарі — це рядки коду, які JavaScript навмисно ігнорує. Коментарі — чудовий спосіб залишити нотатки для себе та інших людей, які пізніше мають зрозуміти, що робить цей код.

Існує два способи написання коментарів у JavaScript:

Використання `//` скаже JavaScript ігнорувати остачу тексту на поточній лінії. Це однолінійний коментар:

```js
// This is an in-line comment.
```

Ви можете створити багатолінійний коментар, починаючи з `/*` і закінчуючи `*/`. Це багатолінійний коментар:

```js
/* This is a
multi-line comment */
```

**ПРИМІТКА:** Під час написання коду ви повинні регулярно додавати коментарі, щоб прояснити призначення частин вашого коду. Хороші коментарі допомагають передати намір вашого коду — як іншим, так і вашому майбутньому «я».

# --instructions--

Спробуйте створити по одному коментарю кожного типу.

# --hints--

Ви повинні створити коментар у стилі `//`, який містить щонайменше п’ять літер.

```js
assert(code.match(/(\/\/)...../g));
```

Ви повинні створити коментар у стилі `/* */`, який містить щонайменше п’ять літер.

```js
assert(code.match(/(\/\*)([^\/]{5,})(?=\*\/)/gm));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
// Fake Comment
/* Another Comment */
```
