---
id: bd7993c9c69feddfaeb7bdef
title: Помножити два десяткові числа на JavaScript
challengeType: 1
forumTopicId: 301173
dashedName: multiply-two-decimals-with-javascript
---

# --description--

У JavaScript ви також можете виконувати обчислення з десятковими числами, так само як і з цілими числами.

Давайте помножимо два десяткові числа, щоб отримати їхній добуток.

# --instructions--

Змініть `0.0`, щоб product дорівнювало `5.0`.

# --hints--

Змінна `product` має дорівнювати `5.0`.

```js
assert(product === 5.0);
```

Ви маєте використати оператор `*`

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 2.0 * 0.0;
```

# --solutions--

```js
const product = 2.0 * 2.5;
```
