---
id: cf1111c1c11feddfaeb3bdef
title: додати два числа за допомогою JavaScript
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` — це тип даних у JavaScript, який представляє числові дані.

Тепер спробуємо додати два числа за допомогою JavaScript.

JavaScript використовує символ `+` як оператор додавання, коли він розміщений між двома числами.

**Приклад:**

```js
const myVar = 5 + 10;
```

`myVar` тепер має значення `15`.

# --instructions--

Змініть `0`, щоб сума дорівнювала `20`.

# --hints--

`sum` має дорівнювати `20`.

```js
assert(sum === 20);
```

Ви маєте використати оператор `+`.

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
