---
id: cf1231c1c11feddfaeb5bdef
title: Помножити два числа за допомогою JavaScript
challengeType: 1
forumTopicId: 18243
dashedName: multiply-two-numbers-with-javascript
---

# --description--

Ми також можемо помножити одне число на інше.

JavaScript використовує символ `*` для множення двох чисел.

**Приклад**

```js
const myVar = 13 * 13;
```

`myVar` матиме значення `169`.

# --instructions--

Змініть `0` так, щоб добуток дорівнював `80`.

# --hints--

Змінна `product` має дорівнювати 80.

```js
assert(product === 80);
```

Ви маєте використати оператор `*`.

```js
assert(/\*/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const product = 8 * 0;
```

# --solutions--

```js
const product = 8 * 10;
```
