---
id: cf1391c1c11feddfaeb4bdef
title: Створення десяткових чисел за допомогою JavaScript
challengeType: 1
forumTopicId: 16826
dashedName: create-decimal-numbers-with-javascript
---

# --description--

Ми також можемо зберігати десяткові числа у змінних. Десяткові числа іноді називають <dfn>числами з плаваючою комою</dfn> або <dfn>флоатами</dfn>.

**Примітка:** коли ви обчислюєте числа, вони обчислюються з кінцевою точністю. Операції з числами з плаваючою комою можуть призводити до результатів, які відрізняються від бажаного. Якщо ви отримуєте один із таких результатів, відкрийте тему на <a href="https://forum.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">форумі freeCodeCamp</a>.

# --instructions--

Створіть змінну `myDecimal` і надайте їй десяткове значення з дробовою частиною (наприклад, `5.7`).

# --hints--

`myDecimal` має бути числом.

```js
assert(typeof myDecimal === 'number');
```

`myDecimal` має містити десяткову крапку

```js
assert(myDecimal % 1 != 0);
```

# --seed--

## --seed-contents--

```js
const ourDecimal = 5.7;

// Only change code below this line

```

# --solutions--

```js
const myDecimal = 9.9;
```
