---
id: 61537bb9b1a29430ac15ad38
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

通过将你的 `styles.css` 文件链接到页面来开始这个项目。

# --hints--

你的 `head` 元素内应包含一个 `link` 元素。

```js
assert.exists(document.querySelector("head>link"));
```

你的 `link` 元素应包含 `rel="stylesheet"` 和 `href="styles.css"`。

```js 
  const link = document.querySelector("head>link"); 
  assert.equal(link?.getAttribute("rel"), "stylesheet"); 
  const href = link?.getAttribute("data-href"); 
  assert.oneOf(href, ["./styles.css", "styles.css"]); 
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
--fcc-editable-region--

--fcc-editable-region--    
  </head>
  <body>
  </body>
</html>
```

```css

```
