---
id: a3566b1109230028080c9345
title: 設計一個加總所有數字的演算法
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

在這個實驗中，你需要設計一個將所有數字加總的演算法。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `sumAll` 的函式（程式），它接受一個包含兩個數字的陣列。
1. `sumAll([n, m])` 應該傳回 `n` 和 `m` 的總和加上它們之間所有數字的總和。最低的數字不一定會先出現。 例如，`sumAll([4,1])` 應該傳回 `10`，因為 `1` 和 `4` 之間（含兩端）所有數字的總和是 `10`。

# --hints--

你應該有一個名為 `sumAll` 的函式（程式）。

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` 應該傳回一個數字。

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` 應該傳回 `10`。

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` 應該傳回 `10`。

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` 應該傳回 `45`。

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` 應該傳回 `45`。

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
