---
id: 66ad8ab945d266383f318cbf
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Sasa unapaswa kuona ujumbe wa kwanza kutoka kwa bot kwenye konsoli.

Ni wakati wa kuongeza ujumbe wa pili kutoka kwa bot.

Ongeza tamko jingine la `console.log` ili kutoa ujumbe `"I am excited to talk to you."` kwenye konsoli.

# --hints--

Unapaswa kuwa na tamko la pili la `console.log()` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

Tamko lako la `console` linapaswa kuwa na ujumbe `"I am excited to talk to you."`. Usisahau alama za nukuu.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
