---
id: 66ad998285df023b23dacdd3
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa, ni wakati wa kuanzisha mabadiliko `botLocation`.

Unapohitaji kutangaza mabadiliko yenye maneno mengi, unaweza kutumia mkataba wa majina wa <dfn>camelCase</dfn>.

Unapotumia `camelCase`, neno la kwanza huwa herufi ndogo kabisa na herufi ya kwanza ya kila neno linalofuata huandikwa kwa herufi kubwa.

Hapa kuna mfano:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

Tangaza na weka thamani ya mfuatano wa herufi `"the universe"` kwa mabadiliko `botLocation` katika mstari mmoja kwa kutumia neno kuu la `let`.

# --hints--

Tumia neno kuu la `let` kutangaza mabadiliko.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Mabadiliko yako yanapaswa kuitwa `botLocation`.

```js
assert.isNotNull(botLocation);
```

Unapaswa kuweka thamani ya mfuatano wa herufi kwa mabadiliko yako ya `botLocation`.

```js
assert.isString(botLocation);
```

Mabadiliko yako ya `botLocation` yanapaswa kupewa thamani ya mfuatano wa herufi `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
