---
id: 66adb844118ba74107ce771f
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Ujumbe unaofuata kutoka kwa bot utahusu mahali bot iko.

Tengeneza kigezo kinachoitwa `botLocationSentence`.

Kisha tumia muunganiko wa mfuatano wa herufi kwa kutumia kiendeshaji `+` kuunganisha mfuatano wa herufi `"I live in "` na kigezo `botLocation` kisha ukifuatiwa na nukta (`.`).

Weka thamani hii katika kigezo `botLocationSentence`.

Kisha, andika thamani ya `botLocationSentence` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo `botLocationSentence`.

```js
assert.isNotNull(botLocationSentence);
```

Kigezo chako `botLocationSentence` kinapaswa kuwa na thamani ya mfuatano wa herufi.

```js
assert.isString(botLocationSentence);
```

Unapaswa kutumia muunganiko wa mfuatano wa herufi kwa kiendeshaji `+` kuunganisha mfuatano wa herufi `"I live in "` na kigezo `botLocation` kisha ukifuatiwa na nukta (`.`). Angalia nafasi vizuri.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

Unapaswa kuwa na tamko la `console` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

Tamko lako la `console` linapaswa kutoa kigezo `botLocationSentence`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
