---
id: 66adc6046acc18453decf577
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Ili kuona jina jipya la utani la roboti, utahitaji kuandika ujumbe mpya kwenye konsoli.

Tengeneza kigezo kipya kinachoitwa `newNicknameGreeting`.

Kisha tumia kuunganisha mfuatano wa herufi kuunganisha mfuatano wa herufi `"I love my nickname but I wish people would call me "` na kigezo `bot` kifuatacho alama ya nukta.

Weka thamani ya matokeo kwenye kigezo `newNicknameGreeting`.

Kisha, andika thamani ya `newNicknameGreeting` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

Kigezo chako cha `newNicknameGreeting` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(newNicknameGreeting);
```

Unapaswa kutumia kuunganisha mfuatano wa herufi kuunganisha mfuatano wa herufi `"I love my nickname but I wish people would call me "` na kigezo `bot` kifuatacho alama ya nukta.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Unapaswa kuwa na tamko la `console` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Unapaswa kuandika thamani ya `newNicknameGreeting` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
