---
id: 66adc5919c1853448f119d11
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa inaonekana kama bot anataka kubadilisha jina lake la utani.

Kwa kutumia mgawo wa thamani, weka mfuatano wa herufi `"awesomeTeacherBot"` katika kigezo `bot`.

# --hints--

Unapaswa kuweka mfuatano wa herufi `"awesomeTeacherBot"` katika kigezo `bot`.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
