---
id: 66ad8294a0ad902f1b31b612
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii utajifunza jinsi ya kufanya kazi na misingi ya JavaScript kwa kujenga roboti ya salamu.

Katika hatua hii ya kwanza, utataka kutoa ujumbe kwenye konsoli kutoka kwa roboti ya salamu.

Kumbuka kwamba ulijifunza kuhusu `console.log()` na mfuatano wa herufi katika mafundisho yaliyopita.

Hapa ni ukumbusho wa jinsi ya kutumia `console.log()` na mfuatano wa herufi:

```js
console.log("Hello, World!");
```

Ongeza tamko la `console.log()` linalotoa mfuatano wa herufi `"Hi there!"` kwenye konsoli. Usisahau alama zako za nukuu kuzunguka ujumbe! 


# --hints--

Unapaswa kuwa na tamko la `console.log()` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

Tamko lako la `console` linapaswa kuwa na ujumbe `"Hi there!"`. Usisahau alama za nukuu.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
