---
id: 66ada3f46945763dd97f43f8
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Awali, uliunda vigezo vya `bot` na `botLocation`. Sasa, utavitumia kutoa ujumbe mpya kwenye konsoli.

Katika mafundisho ya awali, ulijifunza jinsi ya kufanya mfuatano wa herufi kwa kutumia opereta `+` kuunganisha mfuatano wa herufi pamoja kama ifuatavyo:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Kumbuka kuwa unahitaji kuwa makini na nafasi wakati wa kuunganisha mfuatano wa herufi na vigezo.

Unda kigezo kinachoitwa `botIntroduction`.

Kisha tumia mfuatano wa herufi kwa opereta `+` kuunganisha mfuatano wa herufi `"My name is "` ukifuatiwa na kigezo `bot` kisha alama ya nukta (`.`).

Weka thamani hii kwenye kigezo `botIntroduction`.

Kisha, andika thamani ya kigezo `botIntroduction` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

Kigezo chako cha `botIntroduction` kinapaswa kuwa na thamani ya mfuatano wa herufi.

```js
assert.isString(botIntroduction);
```

Unapaswa kutumia mfuatano wa herufi kwa opereta `+` kuunganisha mfuatano wa herufi `"My name is "` na kigezo `bot` kisha alama ya nukta (`.`). Kuwa makini na nafasi. Weka mfuatano huu uliounganishwa kwenye kigezo `botIntroduction`.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

Unapaswa kuwa na tamko la nne la `console.log()` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

Tamko lako la `console` linapaswa kutoa thamani ya kigezo `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
