---
id: 66adcf383276814776aba3ca
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Ifuatayo, tengeneza kigezo kinachoitwa `favoriteSubjectSentence`.

Tumia muunganiko wa mfuatano wa herufi kuunganisha mfuatano wa herufi `"My favorite subject is "` na kigezo `favoriteSubject` kisha ukifuate na nukta.

Weka thamani ya matokeo katika kigezo `favoriteSubjectSentence`.

Kisha, andika thamani ya `favoriteSubjectSentence` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `favoriteSubjectSentence`.

```js
assert.isNotNull(favoriteSubjectSentence);
```

Kigezo chako cha `favoriteSubjectSentence` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(favoriteSubjectSentence);
```

Unapaswa kutumia muunganiko wa mfuatano wa herufi kuunganisha mfuatano wa herufi `"My favorite subject is "` na kigezo `favoriteSubject` kisha ukifuate na nukta. Angalia nafasi vizuri.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

Unapaswa kuwa na tamko la `console` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

Unapaswa kuandika thamani ya `favoriteSubjectSentence` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
