---
id: 66ad984e92ce213a66f7e51c
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa, ni wakati wa kuweka thamani kwa mabadiliko wako wa `bot`.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuweka thamani kwa mabadiliko kama hivi:

```js
variableName = "Here is the value";
```

Kumbuka kwamba kile kilicho upande wa kulia wa opereta ya mgawo `=` ni thamani unayoitoa kwa mabadiliko upande wa kushoto.

Weka mfuatano wa herufi `"teacherBot"` kwa mabadiliko ya `bot`.

# --hints--

Unapaswa kuweka mfuatano wa herufi kwa mabadiliko yako ya `bot`.

```js
assert.isString(bot);
```

Mabadiliko yako ya `bot` yanapaswa kupewa thamani ya mfuatano wa herufi `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
