---
id: 66adc82170acb1464bd348eb
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Ujumbe wa mwisho kutoka kwa bot utazingatia somo la bot anapenda zaidi.

Anza kwa kuunda kigezo kipya kinachoitwa `favoriteSubject` na ukaweke thamani ya mfuatano wa herufi `"Computer Science"`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `favoriteSubject`.

```js
assert.isNotNull(favoriteSubject);
```

Kigezo chako cha `favoriteSubject` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(favoriteSubject);
```

Unapaswa kuweka thamani ya mfuatano wa herufi `"Computer Science"` kwa kigezo cha `favoriteSubject`.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
