---
id: 66ad9c633fa26d3c1475eae3
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa, ni wakati wa kuongeza ujumbe mwingine kutoka kwa bot.

Ongeza tamko jingine la `console` kwenye msimbo ambalo linaandika ujumbe `"Allow me to introduce myself."`.

# --hints--

Unapaswa kuwa na tamko la tatu la `console.log()` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

Tamko lako la `console` linapaswa kuwa na ujumbe `"Allow me to introduce myself."`. Usisahau alama za nukuu.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
