---
id: 66ad8d150264db3926eccfeb
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Katika mafundisho yaliyopita, ulijifunza kuhusu neno kuu la `let` na jinsi ya kutangaza vigezo katika JavaScript.

Hapa ni ukumbusho wa jinsi ya kutangaza kigezo kwa kutumia neno kuu la `let`:

```js
let greeting;
```

Tumia neno kuu la `let` kutangaza kigezo kinachoitwa `bot`.

**KUMBUKUMBU**: Unatumia `let` hapa kwa sababu baadaye katika warsha, utabadilisha thamani ya kigezo cha `bot`.

# --hints--

Unapaswa kuwa na neno kuu la `let` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

Unapaswa kutangaza kigezo kinachoitwa `bot`.

```js
assert.isNotNull(bot);
```

Unapaswa kutumia neno kuu la `let` kutangaza kigezo hicho.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
