---
id: 66adbc8a9793d64250f7e609
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuweka thamani upya kwa vigezo kama hivi:

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

Ukitegemea mgawo wa thamani upya, weka mfuatano wa herufi `"professorBot"` katika kigezo `bot`.

# --hints--

Unapaswa kuweka thamani mfuatano wa herufi `"professorBot"` katika kigezo `bot`.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
