---
id: 66adc42868cab843ccee87d9
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Sasa ni wakati wa kuona thamani mpya ya `bot`.

Anza kwa kuunda kigezo kipya kinachoitwa `nicknameIntroduction`.

Tumia kuunganisha mfuatano wa herufi kuunganishia mfuatano wa herufi `"My nickname is "` na kigezo `bot` kisha ukifuate na nukta (`.`).

Weka thamani ya mfuatano wa herufi uliopatikana katika kigezo `nicknameIntroduction`.

Kisha, andika thamani ya `nicknameIntroduction` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `nicknameIntroduction`.

```js
assert.isNotNull(nicknameIntroduction);
```

Kigezo chako cha `nicknameIntroduction` kinapaswa kuwa mfuatano wa herufi.

```js
assert.isString(nicknameIntroduction);
```

Unapaswa kutumia kuunganisha mfuatano wa herufi kuunganishia mfuatano wa herufi `"My nickname is "` na kigezo cha `bot` kisha ukifuate na nukta (`.`). Angalia nafasi vizuri.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

Unapaswa kuwa na tamko la `console` katika msimbo wako.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

Unapaswa kuandika thamani ya `nicknameIntroduction` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
