---
id: 66fe4f33a2cc9b33f4d5cd9b
title: Costruisci un generatore di permutazioni
challengeType: 26
dashedName: build-a-permutation-generator
---

# --description--

In questo laboratorio costruirai un generatore di permutazioni che prende una stringa e restituisce tutte le possibili permutazioni dei caratteri nella stringa. Per esempio, le permutazioni possibili della stringa `cat` sono `cat`, `cta`, `act`, `atc`, `tac` e `tca`.

Il modo ricorsivo di creare permutazioni di una stringa funziona memorizzando la parte iniziale fissa della stringa (prefisso) e creando permutazioni del resto.

Per esempio, considera la parola `machine`. Il primo giro di creazione delle permutazioni si fa fissando `m` come prefisso della stringa, e poi creando permutazioni del resto della stringa, `achine`.

Per il resto della stringa, le permutazioni continuano allo stesso modo. Una lettera viene aggiunta al prefisso, magari `c`, così il prefisso diventa `mc`. Poi, ciascuna delle permutazioni di `ahine` viene concatenata al prefisso.

Questo continua finché il prefisso contiene tutte le lettere e il resto della stringa è vuoto, il che significa che è stata creata una permutazione.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `permuteString`.
2. La funzione `permuteString` deve prendere un parametro, una stringa, e poi due parametri con valore predefinito: un valore prefisso e un array vuoto per memorizzare e restituire i risultati. Il valore prefisso deve essere una stringa usata per accumulare caratteri e formare una permutazione. La funzione sarà chiamata con un solo argomento, come `permuteString("cat")`.
3. All’interno della funzione, devi verificare se la lunghezza della stringa passata è `0`. Se lo è, aggiungi il prefisso corrente ai risultati e restituisci i risultati.
4. Itera su ogni carattere della stringa di input e per ogni iterazione, rimuovi il carattere corrente dalla stringa e chiama ricorsivamente la funzione `permuteString` con argomenti aggiornati per costruire le permutazioni rimanenti.
5. Devi restituire l’array finale dei risultati.
6. Devi assicurarti che le permutazioni siano uniche rimuovendo i duplicati.

# --hints--

Devi avere una funzione `permuteString`.

```js
assert.isFunction(permuteString);
```

Devi usare la ricorsione nella tua funzione `permuteString`.

```js
assert.match(permuteString.toString(), /permuteString\s*\(/);
```

`permuteString("far")` dovrebbe restituire `[ "far", "fra", "afr", "arf", "rfa", "raf" ]`.

```js
assert.sameMembers(permuteString("far"), [ "far", "fra", "afr", "arf", "rfa", "raf" ]);
```

`permuteString("fcc")` dovrebbe restituire `[ "fcc", "cfc", "ccf" ]`.

```js
assert.sameMembers(permuteString("fcc"), [ "fcc", "cfc", "ccf" ]);
```

`permuteString("p")` dovrebbe restituire `[ "p" ]`.

```js
assert.deepStrictEqual(permuteString("p"), ["p"]);
```

`permuteString("")` dovrebbe restituire `[""]`.

```js
let emptyArr= permuteString("");

// 1 because the empty string is being pushed and is a permutation of itself
assert.lengthOf(emptyArr, 1);

```

`permuteString("walk")` dovrebbe restituire `["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]`.
`.

```js
assert.sameMembers(permuteString("walk"), ["walk", "wakl", "wlak", "wlka", "wkla", "wkal", "awlk", "awkl", "alwk", "alkw", "aklw", "akwl", "lawk", "lakw", "lwak", "lwka", "lkaw", "lkwa", "kawl", "kalw", "kwal", "kwla", "klaw", "klwa"]);
```

`permuteString` dovrebbe restituire i risultati corretti.

```js
// tests to prevent hard coding
assert.sameMembers(permuteString("road"), [ "road",
  "roda",
  "raod",
  "rado",
  "rdoa",
  "rdao",
  "orad",
  "orda",
  "oard",
  "oadr",
  "odra",
  "odar",
  "arod",
  "ardo",
  "aord",
  "aodr",
  "adro",
  "ador",
  "droa",
  "drao",
  "dora",
  "doar",
  "daro",
  "daor" ]);

assert.sameMembers(permuteString("fog"), [ "fog",
  "fgo",
  "ofg",
  "ogf",
  "gfo",
  "gof" ]);

assert.sameMembers(permuteString("bird"), [ "bird",
  "bidr",
  "brid",
  "brdi",
  "bdir",
  "bdri",
  "ibrd",
  "ibdr",
  "irbd",
  "irdb",
  "idbr",
  "idrb",
  "rbid",
  "rbdi",
  "ribd",
  "ridb",
  "rdbi",
  "rdib",
  "dbir",
  "dbri",
  "dibr",
  "dirb",
  "drbi",
  "drib" ]);

```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const permuteString = (str, prefix = "", result = []) => {
  if (str.length === 0) {
    result.push(prefix);
    return result;
  }

  // Sort the string initially to group duplicates
  const sortedStr = str.split("").sort().join("");

  for (let i = 0; i < sortedStr.length; i++) {
    if (i > 0 && sortedStr[i] === sortedStr[i - 1]) {
      // Skip duplicate characters
      continue;
    }

    const rem = sortedStr.slice(0, i) + sortedStr.slice(i + 1);
    permuteString(rem, prefix + sortedStr[i], result);
  }

  return result;
};
```
