---
id: 66bb3e20d3dc5b6d0a21f5dd
title: costruisci un contatore in tempo reale
challengeType: 25
dashedName: build-a-real-time-counter
demoType: onClick
---

# --description--

In questo laboratorio costruirai un contatore di caratteri in tempo reale usando JavaScript. Il contatore di caratteri mostrerà il numero di caratteri inseriti in un elemento `textarea`. Il contatore si aggiornerà in tempo reale mentre l’utente digita nel `textarea`.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `textarea` con l’`id` uguale a `text-input`.
2. Deve esserci un elemento `p` con l’`id` uguale a `char-count`. Dovrebbe contenere inizialmente il testo `Character Count: 0/50`. Questo testo segnaposto deve essere sostituito con il conteggio attuale dei caratteri nell’elemento `textarea`.
3. Quando l’elemento `#text-input` contiene il testo `hello coder`, l’elemento `#char-count` deve contenere il testo `"Character Count: 11/50"`.
4. Quando l’elemento `#text-input` contiene il testo `i am a programmer`, l’elemento `#char-count` deve contenere il testo `"Character Count: 17/50"`.
5. Quando l’elemento `#text-input` contiene il testo `hello world`, l’elemento `#char-count` deve contenere il testo `"Character Count: 11/50"`.
6. Quando l’elemento `#text-input` contiene il testo `I am learning a new language and it is called c++.`, l’elemento `#char-count` deve contenere il testo `"Character Count: 50/50"`. **NOTA**: Anche se l’attributo `maxlength` raggiungerebbe la funzionalità corretta, questo laboratorio richiede una soluzione solo in JavaScript.  
7. L’utente non deve poter inserire più di `50` caratteri nell’elemento `textarea`. Quando il conteggio dei caratteri raggiunge `50`, qualsiasi input extra deve essere automaticamente tagliato e il testo `Character Count: 50/50` deve essere mostrato in `red`.

**Nota:** Assicurati di collegare il tuo foglio di stile e il file JavaScript nel tuo HTML.

# --hints--

Dovresti avere un elemento `textarea` con l’`id` uguale a `text-input`. 

```js
const textInput = document.querySelector('textarea');
assert.equal(textInput.id, 'text-input');
```

Dovresti avere un elemento `p` con l’`id` uguale a `char-count`.

```js
assert.exists("p#char-count");
```

L’elemento `#char-count` dovrebbe contenere inizialmente il testo `Character Count: 0/50`.

```js
const charCount = document.querySelector('p#char-count')?.innerText;
assert.equal(charCount, 'Character Count: 0/50');
```

Quando l’elemento `#text-input` contiene il testo `hello coder`, l’elemento `#char-count` deve contenere il testo `"Character Count: 11/50"`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "hello coder";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(charCount.innerText.trim(), 'Character Count: 11/50');
```

Quando il conteggio dei caratteri è `50`, il testo deve essere mostrato in `red`.

```js
const textInput = document.getElementById('text-input');
const charCount = document.getElementById('char-count');
textInput.value = "I am learning a new language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
const computedStyle = window.getComputedStyle(charCount);
assert.oneOf(computedStyle.color, [
  "red",
  "rgb(255, 0, 0)",
  "rgb(100%, 0%, 0%)",
  "rgba(255, 0, 0, 1)",
  "rgba(100%, 0%, 0%, 1)",
  "hsl(0, 100%, 50%)",
  "hsla(0, 100%, 50%, 1)"
]);
```

Se il conteggio dei caratteri è maggiore o uguale a `50`, l’input extra deve essere tagliato.

```js
const textInput = document.getElementById('text-input');
textInput.value = "I am learning a new programming language and it is called c++.";
textInput.dispatchEvent(new Event('input'));
textInput.dispatchEvent(new Event('change'));
assert.strictEqual(textInput.value.length, 50);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real Time Counter</title>
</head>

<body>

</body>


</html>
```

```css

```

```js
    
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Real-time Character Counter</title>
    <link rel="stylesheet" href="styles.css">

<body>
    <div class="container">
        <h1>Real-Time Character Counter</h1>
        <textarea id="text-input" placeholder="Type something..."></textarea>
        <p id="char-count">Character Count: 0/50</p>
    </div>

<script src="script.js"></script>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    font-family: Arial, sans-serif;
    background-color: #c0c1ec;
    margin: 0;
}

h1{
    background-color: #f2f48b;
    padding-left: 10px;
    padding-right: 10px;
}
.container {
    text-align: center;
}

textarea {
    width: 300px;
    height: 100px;
    font-size: 18px;
    padding: 10px;
    border: 1px solid #ccc;
    border-radius: 5px;
    resize: none;
    display: block;
    margin: 0 auto;
    border: 1px solid slategray;
    box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

textarea:focus {
    outline: none;
    border-color: #f2f48b;
}

#char-count {
    margin-top: 10px;
    font-size: 18px;
    color: darkslategray;
    font-weight: bold;
}
```

```js
document.getElementById("text-input").addEventListener("input", function () {
  let textInput = document.getElementById("text-input");
  let charCount = textInput.value.length;

  if (charCount > 50) {
    textInput.value = textInput.value.substring(0, 50);
    charCount = 50;
  }

  let charCountDisplay = document.getElementById("char-count");
  charCountDisplay.innerHTML = `Character Count: ${charCount}/50`;

  if (charCount === 50) {
    charCountDisplay.style.color = "red";
  } 

});
```
