---
id: 5f33294a6af5e9188dbdb8f3
title: 1단계
challengeType: 0
dashedName: step-1
---

# --description--

이 워크숍에서는 카페 메뉴를 만들면서 CSS(Cascading Style Sheets)의 기본을 연습합니다.

메뉴 내용을 추가하는 것부터 시작하세요. 기존 `main` 요소 안에 `body` 요소를 추가하세요. 이 요소에는 나중에 카페에서 제공하는 커피와 디저트의 가격 정보가 들어갑니다.

# --hints--

코드에는 `<main>` 여는 태그가 있어야 합니다.

```js
assert.match(code, /<main>/i);
```

코드에는 `</main>` 닫는 태그가 있어야 합니다.

```js
assert.match(code, /<\/main>/i);
```

`body` 요소는 변경하지 마세요. 닫는 태그를 실수로 삭제하지 않도록 주의하세요.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

`main` 태그는 `body` 태그 안에 있어야 합니다.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
