---
id: 5f332a88dc25a0fd25c7687a
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

카페 이름은 `CAMPER CAFE`입니다. 그래서 `h1` 요소 안에 `main` 요소를 추가하세요. 카페 이름을 대문자로 작성해 눈에 띄게 만드세요.

# --hints--

열린 `<h1>` 태그가 있어야 합니다.

```js
assert.match(code, /<h1>/i);
```

닫힌 `</h1>` 태그가 있어야 합니다.

```js
assert.match(code, /<\/h1>/i);
```

기존 `main` 요소는 변경하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

`h1` 요소는 `main` 요소 안에 중첩되어야 합니다.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

`h1` 요소에는 `CAMPER CAFE` 텍스트가 있어야 합니다.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
