---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

지금까지는 콘텐츠의 표현과 모양을 제한적으로 제어했습니다. 이를 바꾸려면 `style` 요소 안에 `head` 요소를 추가하세요.

# --hints--

코드에 여는 `<style>` 태그가 있어야 합니다.

```js
assert.match(code, /<style\s*>/i);
```

코드에 닫는 `</style>` 태그가 있어야 합니다.

```js
assert.match(code, /<\/style\s*>/);
```

`style` 요소는 `head` 요소 안에 중첩되어야 합니다.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
