---
id: 5f3477cbcb6ba47918c1da92
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

모바일에서 페이지 스타일이 데스크톱이나 노트북과 비슷하게 보이도록 하려면, 특별한 `meta` 속성을 가진 `content` 요소를 추가해야 합니다.

이전 수업에서 뷰포트 `meta` 요소에 대해 배웠습니다.

예시는 다음과 같습니다:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

코드에는 `meta` 요소가 두 개 있어야 합니다.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

`meta` 요소에는 값이 `name`인 `viewport` 속성이 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

`meta` 요소에는 값이 `content`인 `width=device-width, initial-scale=1.0` 속성이 있어야 합니다.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
