---
id: 5f344f9c805cd193c33d829c
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이전 수업에서, 다음과 같이 `CSS` 속성과 값을 추가하는 방법을 배웠습니다:

```css
element {
 property: value;
}
```

`h1` 요소의 내용을 가운데 정렬하려면 `text-align` 속성을 `center` 값으로 설정하세요.

# --hints--

`h1` 요소에 `style` 선택자가 있어야 합니다.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

`text-align` 속성은 `center` 값을 가져야 합니다.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

`h1` 선택자는 `text-align` 속성을 `center` 값으로 설정해야 합니다.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
