---
id: 5f3477ae34c1239cafe128be
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

지금 세 개의 타입 선택자가 같은 스타일을 가지고 있습니다. 선택자 목록을 만들어 여러 요소에 같은 스타일 그룹을 추가할 수 있습니다. 각 선택자는 쉼표로 구분합니다:

```css
selector1, selector2 {
  property: value;
}
```

기존의 세 개 타입 선택자를 삭제하고 `h1`, `h2`, `p` 요소의 텍스트를 가운데 정렬하는 하나의 선택자 목록으로 바꾸세요.

# --hints--

세 요소 모두에 대해 하나의 타입 선택자 `h1, h2, p`를 사용하세요. 그 순서를 꼭 지키세요.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

`style` 요소에는 선택자가 하나만 있어야 합니다.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

