---
id: 5f344fbc22624a2976425065
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

`h2` 요소 안에 `section` 요소를 만들고 텍스트 `Coffee`를 넣으세요.

# --hints--

열리는 `<h2>` 태그가 있어야 합니다.

```js
assert.match(code, /<h2\s*>/i);
```

닫는 `</h2>` 태그가 있어야 합니다.

```js
assert.match(code, /<\/h2\s*>/i);
```

기존 `section` 요소를 변경하지 마세요. 닫는 태그를 삭제하지 않았는지 확인하세요.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

`h2` 요소는 `section` 요소 안에 있어야 합니다.

```js
const h2 = document.querySelector('h2');
assert.equal(h2.parentElement.tagName, 'SECTION');
```

`h2` 요소는 텍스트 `Coffee`를 가져야 합니다.

```js
const h2 = document.querySelector('h2');
assert.equal(h2?.innerText.trim(), 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
--fcc-editable-region--
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
