---
id: 5f356ed69db0a491745e2bb6
title: Step 23
challengeType: 0
dashedName: step-23
---

# --description--

클래스의 스타일을 `div` 요소에 적용하려면 `id` 속성을 제거하고 `class` 요소의 여는 태그에 `div` 속성을 추가하세요. 클래스 값은 `menu`로 설정해야 합니다.

# --hints--

`div`는 여전히 렌더링되어야 합니다. `<div>` 태그가 잘못 작성되지 않았는지 확인하세요.

```js
assert.lengthOf(document.querySelectorAll('div'), 1);
```

`div` 요소는 `menu` 클래스를 가져야 합니다.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

`div` 요소는 더 이상 `menu` id를 가져서는 안 됩니다.

```js
assert.lengthOf(document.querySelectorAll('div#menu'), 0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

