---
id: 5f35e5c4321f818cdc4bed30
title: Step 25
challengeType: 0
dashedName: step-25
---

# --description--

이제 화면이 잘 보이니 메뉴 항목을 추가할 차례입니다.

`article` 제목 아래에 빈 `Coffee` 요소를 추가하세요. 이 요소에는 현재 제공하는 각 커피의 맛과 가격이 들어갑니다.

# --hints--

열린 `<article>` 태그가 있어야 합니다.

```js
assert.match(code,/<article>/i);
```

닫힌 `</article>` 태그가 있어야 합니다.

```js
assert.match(code, /<\/article>/i);
```

기존 `h2` 요소는 변경하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('h2'), 1);
```

`article` 요소는 `h2` 요소 뒤에 와야 합니다.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

