---
id: 5f35e5c44359872a137bd98f
title: Step 24
challengeType: 0
dashedName: step-24
---

# --description--

카페의 주요 판매 제품이 커피이므로, 페이지 배경으로 커피 원두 이미지를 사용할 수 있습니다.

`body` 타입 선택자 안에 있는 주석과 그 내용을 제거하세요. 그다음 `background-image` 속성을 추가하고 값을 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`로 설정하세요.

# --hints--

주석 처리된 `background-color` 속성을 제거해야 합니다.

```js
assert.notMatch(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```

`body` 선택자에는 주석이 없어야 합니다.

```js
assert.notMatch(code, /body\s*{\s*\/\*/i);
```

`background-image` 속성을 `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`로 설정해야 합니다.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert.isTrue(hasBackground)
```

`body` 요소에 커피 원두 배경 이미지가 있어야 합니다.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert.equal(bodyBackground, `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  /*
  background-color: burlywood;
  */
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

