---
id: 5f3cade9993019e26313fa8e
title: Step 38
challengeType: 0
dashedName: step-38
---

# --description--

작동하는 것을 알았으니, 나머지 `article`와 `p` 요소들도 첫 번째 세트와 일치하도록 변경하세요. 다른 `item` 요소에 `article` 클래스를 추가하는 것부터 시작하세요.

# --hints--

`article` 요소는 다섯 개만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

`.item` 요소는 다섯 개만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.item'),5);
```

`.item` 요소는 `article` 요소여야 합니다.


```js
const articles = document.querySelectorAll('article');
const items = document.querySelectorAll('.item');
assert.equal(articles?.[0], items?.[0]);
assert.equal(articles?.[1], items?.[1]);
assert.equal(articles?.[2], items?.[2]);
assert.equal(articles?.[3], items?.[3]);
assert.equal(articles?.[4], items?.[4]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
```

