---
id: 5f3c866de7a5b784048f94b1
title: Step 32
challengeType: 0
dashedName: step-32
---

# --description--

이것이 원하는 바와 비슷하지만, 이제 맛과 가격이 같은 줄에 있으면 좋겠습니다. `p` 요소는 <dfn>블록 레벨</dfn> 요소라서 부모 요소의 전체 너비를 차지합니다.

같은 줄에 배치하려면 `p` 요소에 스타일을 적용해 <dfn>인라인</dfn> 요소처럼 동작하게 만들어야 합니다.

그렇게 하려면 `class` 제목 아래 첫 번째 `item` 요소에 `article` 속성을 값 `Coffee` 으로 추가하세요.

# --hints--

`item` 요소에 `article` 클래스를 적용하세요.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

`item` 클래스 요소는 하나만 있어야 합니다.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

첫 번째 `article` 요소에 `item` 클래스를 적용하세요.

```js
assert.equal(document.querySelector('article')?.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          <article>
--fcc-editable-region--
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
