---
id: 5f3c866d28d7ad0de6470505
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

맛과 가격이 현재 각각의 `p` 요소와 함께 위아래로 쌓여 있고 가운데 정렬되어 있습니다. 맛은 왼쪽에, 가격은 오른쪽에 있으면 좋겠습니다.

`flavor` 요소에 `French Vanilla` 클래스 이름을 추가하세요.

# --hints--

`flavor` 요소에 `p` 클래스를 추가해야 합니다.

```js
const el = document.querySelector("p.flavor");
assert.exists(el);
```

`flavor` 클래스는 하나의 요소에만 있어야 합니다.

```js
const elements = document.querySelectorAll('.flavor');
assert.lengthOf(elements, 1);
```

`flavor` 클래스는 `p` 텍스트가 있는 `French Vanilla` 요소에 있어야 합니다.

```js
const el = document.querySelector(".flavor");
assert.equal(el?.innerText.trim(), "French Vanilla");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            <p>French Vanilla</p>
--fcc-editable-region--
            <p>3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
