---
id: 5f3ef6e0a81099d9a697b550
title: 61단계
challengeType: 0
dashedName: step-61
---

# --description--

`address` 안에 `p` 요소를 추가하세요. 그런 다음 `a` 안에 앵커(`p`) 요소를 중첩하여 `https://www.freecodecamp.org`로 링크하고 텍스트는 `Visit our website`로 만드세요.

링크가 새 탭에서 열리도록 `target` 속성을 값 `_blank`로 추가하세요.

# --hints--

기존 `footer` 요소는 수정하지 마세요.

```js
assert.lengthOf(document.querySelectorAll("footer"), 1);
```

기존 `address` 요소는 수정하지 마세요.

```js
assert.lengthOf(document.querySelectorAll("address"), 1);
```

새로운 `p` 요소는 `address` 요소 안에 중첩되어야 합니다. `p` 요소가 정확히 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p"), 1);
assert.lengthOf(document.querySelectorAll("footer p"), 1);
```

새로운 `a` 요소는 새로운 `p` 요소 안에 중첩되어야 합니다. `a` 요소가 정확히 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p > a"), 1);
assert.lengthOf(document.querySelectorAll("footer a"), 1);
```

새로운 `a` 요소는 텍스트 `Visit our website`를 가져야 합니다.

```js
assert.equal(document.querySelector("footer > address > p > a")?.innerText.trim(), "Visit our website");
```

새로운 `a` 요소는 `https://www.freecodecamp.org`로 링크해야 합니다. `a` 요소는 링크를 만들기 위해 `href` 속성을 사용한다는 점을 기억하세요.

```js
assert.equal(document.querySelector("footer > address > p > a")?.href, "https://www.freecodecamp.org/");
```

새로운 `a` 요소는 `target` 속성을 `_blank`로 설정해야 합니다.

```js
assert.equal(document.querySelector("footer > address > p > a")?.target, "_blank");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <footer>
        <address>
--fcc-editable-region--
          
--fcc-editable-region--
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```
