---
id: 5f45a276c093334f0f6e9df4
title: Step 70
challengeType: 0
dashedName: step-70
---

# --description--

메뉴 항목과 가격에 집중하면 각 줄 사이에 꽤 큰 간격이 있습니다.

클래스 이름이 `p`인 요소 안에 중첩된 모든 `item` 요소를 대상으로 하는 기존 선택자를 사용하고, 그들의 위쪽과 아래쪽 여백을 `5px`로 설정하세요.

# --hints--

`margin-top` 속성을 `5px`로 설정해야 합니다.

```js
const hasMarginTop = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-top'] === '5px');
assert.isTrue(hasMarginTop);
```

`margin-bottom` 속성을 `5px`로 설정해야 합니다.

```js
const hasMarginBottom = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-bottom'] === '5px');
assert.isTrue(hasMarginBottom);
```

기존 `.item p` 선택자를 사용해야 합니다.

```js
const hasOneSelector = new __helpers.CSSHelp(document).getStyleDeclarations('.item p');
assert.lengthOf(hasOneSelector, 1);
```

`p` 요소 안에 중첩된 `.item` 요소는 `margin-top`이 `5px`이어야 합니다.

```js
const itemPMarginTop = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-top');
assert.equal(itemPMarginTop, '5px');
```

`p` 요소 안에 중첩된 `.item` 요소는 `margin-bottom`가 `5px`이어야 합니다.

```js
const itemPMarginBottom = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-bottom');
assert.equal(itemPMarginBottom, '5px');
```

`p` 요소 안에 중첩된 `.item` 요소는 `margin-left`이 설정되어 있으면 안 됩니다.

```js
const itemPMarginLeft = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-left');
assert.isEmpty(itemPMarginLeft);
```

`p` 요소 안에 중첩된 `.item` 요소는 `margin-right`이 설정되어 있으면 안 됩니다.

```js
const itemPMarginRight = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('margin-right');
assert.isEmpty(itemPMarginRight);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <hr>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <hr>
      <footer>
        <address>
          <p>
            <a href="https://www.freecodecamp.org" target="_blank">Visit our website</a>
          </p>
          <p>123 Free Code Camp Drive</p>
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
  padding: 20px;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

hr {
  height: 2px;
  background-color: brown;
  border-color: brown;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
--fcc-editable-region--
  
--fcc-editable-region--
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

