---
id: 5f716ad029ee4053c7027a7a
title: Step 45
challengeType: 0
dashedName: step-45
---

# --description--

`p` 요소 안에 두 개의 `article` 요소를 중첩하세요. 첫 번째 요소의 텍스트는 `Donut`여야 하고, 두 번째 요소의 텍스트는 `1.50`이어야 합니다. 두 요소를 같은 줄에 두되, 사이에 공백이 없도록 하세요.

# --hints--

기존 `article` 요소는 변경하지 마세요.

```js
assert.lengthOf(document.querySelectorAll('article'), 6);
```

새로운 `article` 요소 안에는 두 개의 `p` 요소가 있어야 합니다.

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraphs = newArticle?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphs, 2);
```

첫 번째 `p` 요소의 텍스트는 `Donut`이어야 합니다.

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraph = newArticle?.querySelector(`:scope ${'p'}`);
assert.match(paragraph?.innerText.trim(), /Donut/i);
```

두 번째 `p` 요소의 텍스트는 `1.50`이어야 합니다.

```js
const newArticle = [...document.querySelectorAll('article')].at(-1);
const paragraphs = newArticle?.querySelectorAll(`:scope ${'p'}`);
assert.match(paragraphs?.[1]?.innerText.trim(), /1\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
--fcc-editable-region--
            
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```

