---
id: 5f7692f7c5b3ce22a57788b6
title: Step 40
challengeType: 0
dashedName: step-40
---

# --description--

스타일링을 완료하려면, 나머지 모든 `flavor` 요소에 해당하는 클래스 이름 `price`와 `p`을 추가하세요.

# --hints--

`.flavor` 요소가 다섯 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.flavor'), 5);
```

`.price` 요소가 다섯 개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.price'), 5);
```

`.flavor` 요소는 텍스트 `p`, `French Vanilla`, `Caramel Macchiato`, `Pumpkin Spice`, `Hazelnut`을 가진 `Mocha` 요소여야 합니다.

```js
const p = document.querySelectorAll('p');
const flavor = document.querySelectorAll('.flavor');
assert.equal(p?.[1], flavor?.[0]);
assert.equal(p?.[3], flavor?.[1]);
assert.equal(p?.[5], flavor?.[2]);
assert.equal(p?.[7], flavor?.[3]);
assert.equal(p?.[9], flavor?.[4]);
```

`.price` 요소는 텍스트 `p`, `3.00`, `3.75`, `3.50`, `4.00`을 가진 `4.50` 요소여야 합니다.

```js
const p = document.querySelectorAll('p');
const price = document.querySelectorAll('.price');
assert.equal(p?.[2], price?.[0]);
assert.equal(p?.[4], price?.[1]);
assert.equal(p?.[6], price?.[2]);
assert.equal(p?.[8], price?.[3]);
assert.equal(p?.[10], price?.[4]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article class="item">
            <p>Caramel Macchiato</p><p>3.75</p>
          </article>
          <article class="item">
            <p>Pumpkin Spice</p><p>3.50</p>
          </article>
          <article class="item">
            <p>Hazelnut</p><p>4.00</p>
          </article>
          <article class="item">
            <p>Mocha</p><p>4.50</p>
          </article>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
```

