---
id: 619665c9abd72906f3ad30f9
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

你將會建造一隻快樂的 Flappy Penguin，並在此過程中進一步練習 CSS 變形和動畫。

從將你的樣式表聯結到頁面開始。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
assert.exists(document.querySelector('link'));
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
