---
id: 6551f3bfed58796ebb1268dc
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

每个字符串都可以被一个数字索引引用。 索引数开始于零。 所以字符串的第一个字符有一个 `0` 的索引。 例如，在字符串 `'Hello World'`, `'H'` 位于索引 `0`, `'e'` 位于索引 `1`等。

字符串的每个字符都可以使用括号符号访问。 您需要在方括号前写上变量名称并在方括号中添加字符索引：

```py
text = 'Hello World'
r = text[8]
```

现在，不是打印 `text`，而是只打印索引 `6` 号位的字符。

# --hints--

You should still call the `print()` function.

```js
assert.match(code, /print\s*\(.*\)/)
```

You should pass `text[6]` to the `print()` function by including it between the parentheses. 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
