---
id: 6551f6f9ec6db270eb83f3db
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

您可以在终端上看到数字 `11` 被打印出来，因为 `'Hello World'` 包含11个字符。

另一个有用的内置函数是 `type()`, 它返回一个变量的数据类型。 修改您的代码`print()`来调用并打印`text`的数据类型.

# --hints--

You should call the `type()` function.

```js
assert.match(code, /type\s*\(.*\)/)
```

You should pass `text` to the `type()` function by including it between the parentheses.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

You should pass `type(text)` to the `print()` function by including it between the parentheses. 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

You should not have `print(len(text))` in your code.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
