---
id: 6551f4259d66026ee28cb205
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

您也可以从字符串结尾处开始访问字符。 最后一个字符的索引为 `-1`, 倒数第二个是 `-2` ，以此类推。

现在修改您现有的 `print()` 来调用并打印你字符串里的最后一个字符。

# --hints--

You should still call the `print()` function.

```js
assert.match(code, /print\s*\(.*\)/)
```

You should pass `text[-1]` to the `print()` function by including it between the parentheses. Pay attention to place the function call at the beginning of the line.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

您的代码中包含无效的句法和/或无效的缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
