---
id: 65521203d7165c7b84b22ad4
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

A loop allows you to systematically go through a sequence of elements and execute actions on each one.

In this case, you'll employ a `for` loop. 以下是迭代 `text` 的方法：

```py
for i in text:
```

`for` 是表示循环类型的关键字。 `i` 是一个变量，它按顺序获取 `text` 中元素的值。 该语句以冒号 `:` 结尾。

Below the line where you declared `alphabet`, write a `for` loop to iterate over `text`. 使用 `i` 作为循环变量。

这样做时，终端出现错误。 你们将在下一步骤中学习这个问题。

# --hints--

你应该使用 `for` 关键字来创建循环。 确保将 `for` 关键字放在行的开头，并在关键字后面留一个空格。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

你应该在 `for` 关键字后面写入 `i` 变量

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

你应该在 `for i` 之后写入 `in` 关键字。 确保在 `in` 关键字周围留一个空格

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

你应该在 `for i in` 之后写入 `text`。 不要忘记添加最后的 `:`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

你的 `for` 循环应放置在代码 `alphabet = 'abcdefghijklmnopqrstuvwxyz'` 行下方。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
