---
id: 655208d59b131e7816f18c96
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

The `print()` function gives you only an output in the console, but functions and methods can have a return value that you can use in your code.

Now assign `alphabet.find(text[0])` to a variable named `index`. In this way, `index` will store the value returned by `alphabet.find(text[0])`.

# --hints--

You should declare a variable named `index`. 请注意，需把变量的名称放在该行的开头位置。

```js
assert.match(code, /^index\s*=/m)
```

You should assign the value returned by `alphabet.find(text[0])` to your `index` variable.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

您的代码中包含无效的句法和/或无效的缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
