---
id: 6552008c0d9d9075cbec9772
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

你将使用 `.find()` 方法来查找消息中每个字母在字母表中的位置。 方法类似于函数，但它属于对象。

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Above, the `.find()` method is *called on* `sentence` (the string to search in), and `'r'` (the character to locate) is passed as the argument. The `sentence.find('r')` call will return `4`, which is the index of the first occurrence of `'r'` in `sentence`.

At the end of your code, call `.find()` on `alphabet` and pass `'z'` as the argument to the method.

# --hints--

你应该调用 `.find()` 方法。

```js
assert.match(code, /find\s*\(.*\)/)
```

你应该对 `alphabet` 变量调用 `.find()` 方法。

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


你应该对 `alphabet` 变量调用 `.find()` 方法，并将 `'z'` 传递给该方法。 注意将函数调用放在行首。

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


你的代码中包含无效的句法和/或无效的缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
