---
id: 6555231eac4d9f19bd3d44b5
title: 步骤 83
challengeType: 20
dashedName: step-83
---

# --description--

`pass` 关键字可以用作未来代码的占位符。 It does not have any effect in your code but it can save you from errors you would get in case of incomplete code:

```py
def foo():
    pass
```

使用 `1` 调用 `vigenere` 进行加密，使用 `-1` 进行解密是可以的，但可能有点神秘。 创建一个名为 `encrypt` 的新函数，该函数采用 `message` 和 `key` 参数，并使用 `pass` 填充函数体。


# --hints--

You should declare a new function called `encrypt`. 请记得在函数体内使用 `pass` 关键字。

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Your `encrypt` function should take `message` and `key` as the parameters. 记住参数传递的顺序是很重要的。

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]

    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
