---
id: 659530ca823d6c47f61daa8a
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

Print your `shifted` variable.

# --hints--

You should print your `shifted` variable. 注意将函数调用放在行首。

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
