---
id: 659db3e3670d3e712be82593
title: 步骤 22
challengeType: 20
dashedName: step-22
---

# --description--

重复在字母表中定位字母并确定 `text` 中每个字符的移位字母的过程可能很乏味。 值得庆幸的是，你可以使用循环来简化它。

现在，移除所有在声明 `alphabet` 变量之后的代码行。

# --hints--

你的代码中仍应包含 `text = 'Hello World'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

你的代码中仍应包含 `shift = 3`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

你的代码中仍应包含 `alphabet = 'abcdefghijklmnopqrstuvwxyz'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

你应该删除 `index` 变量及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

你的代码中不应该有 `print(index)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

你应该删除 `shifted` 变量及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

你的代码中不应该有 `print(shifted)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

你的代码包含无效的语法和/或无效缩进。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
