---
id: 69a56b5069ca99f7317e6e19
title: 建置貨物清單驗證器
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

在這個實驗中，你將使用 JavaScript 來標準化並驗證貨物清單。貨物清單是一份通常列出運輸貨物（例如，透過船舶或火車運輸）的文件，並包含這些貨物的詳細資訊。

每個貨物清單將以物件形式表示，具有以下屬性：

- `containerId`：一個正整數，用來識別相關的貨物容器。
- `destination`：一個非空字串（去除空白後），表示貨物的標的目的地。
- `weight`：一個正數，代表貨物的重量。
- `unit`：一個字串，描述貨物重量屬性的單位（要麼是 `"kg"` 表示公斤，要麼是 `"lb"` 表示磅）。
- `hazmat`：一個布林值，表示是否需要危險物品處理。

貨物清單物件範例：

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**目標：** 完成以下使用者故事並通過所有測試以完成此實驗。

**使用者故事：**

1. 你應該實作一個名為 `normalizeUnits` 的函式，帶有一個 `manifest` 參數。
    - 該函式不得修改原始的清單物件，且必須始終傳回一個新的物件，其中 `weight` 被標準化為公斤，且 `unit` 設為 `"kg"`。
    - 如果清單物件的重量是以磅（`unit: "lb"`）表示，函式應使用近似換算 1 lb = 0.45 kg 將 `weight` 轉換為公斤，並相應更新 `unit`。
    - 如果重量已經是以公斤（`unit: "kg"`）表示，則 `weight` 和 `unit` 應保持不變。
2. 你應該實作一個名為 `validateManifest` 的函式，帶有一個 `manifest` 參數。
    - 該函式不得修改原始的清單物件，且必須始終傳回一個新的物件。
    - 如果輸入的清單有效（沒有遺漏或無效的屬性），函式應傳回一個空物件。
    - 如果輸入的清單無效，函式應傳回一個物件，包含每個遺漏或無效屬性的條目。遺漏的屬性值應為 `"Missing"`，無效的屬性值應為 `"Invalid"`。

    輸入物件遺漏 `destination` 屬性且 `weight` 屬性無效時的範例返回值：

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. 你應該實作一個名為 `processManifest` 的函式，帶有一個 `manifest` 參數。該函式應該記錄：
    - 如果清單物件有效，先記錄 `Validation success: ${containerId}`，然後記錄該清單的公斤重量，格式為 `Total weight: ${weight} kg`。此轉換應使用 `normalizeUnits()`。
    - 如果清單物件無效，先記錄 `Validation error: ${containerId}`，然後記錄呼叫 `validateManifest()` 並傳入該清單物件所返回的物件。
    **注意：** 這兩種情況都應該有兩次 `console.log()` 呼叫。

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

你應該有一個名為 `normalizeUnits` 的函式，帶有一個 `manifest` 參數。

```js
assert.isFunction(normalizeUnits);
const funcString = normalizeUnits.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

呼叫 `normalizeUnits()` 並傳入 `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }` 應該傳回新的物件 `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }`，且不會修改原始輸入。

```js
const testObj = {
  containerId: 68,
  destination: "Salinas", 
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

你的 `normalizeUnits` 函式應該傳回輸入清單物件的複本，並將其 `weight` 標準化為公斤，且將 `unit` 設為 `"kg"`。使用近似換算 `1 lb = 0.45 kg` 來進行重量轉換。

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits should return an object");
  assert.notStrictEqual(result, obj, "normalizeUnits must return a new object, not modify the original");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits did not return the expected normalized object for manifest at index ${i}`
  );
});
```

你的 `normalizeUnits` 函式應該傳回輸入 `manifest` 物件的新複本，且不會修改原始物件。

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

你應該有一個名為 `validateManifest` 的函式，帶有一個 `manifest` 參數。

```js
assert.isFunction(validateManifest);
const funcString = validateManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

呼叫 `validateManifest()` 並傳入 `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }` 應該傳回新的物件 `{}`。

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz", 
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);

assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated");
```

如果輸入的清單物件有效，你的 `validateManifest` 函式應該傳回空物件 `{}`。

```js
const solution = {}
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

呼叫 `validateManifest()` 並傳入 `{}` 應該傳回新的物件 `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`，且不會修改原始輸入。

```js
const testObj = {
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original")
```

呼叫 `validateManifest()` 並傳入 `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }` 應該傳回新的物件 `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }`，且不會修改原始輸入。

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated");
```

呼叫 `validateManifest()` 並傳入 `{ containerId: -2 }` 應該傳回新的物件 `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`，且不會修改原始輸入。

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

呼叫 `validateManifest()` 並傳入 `{ containerId: 3.50 }` 應該傳回新的物件 `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }`，且不會修改原始輸入。你可以使用 `Number.isInteger()` 來驗證整數值。

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

呼叫 `validateManifest()` 並傳入 `{ destination: "  " }` 應該傳回新的物件 `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }`，且不會修改原始輸入。你可以使用 `String.trim()` 來移除字串的空白。

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

呼叫 `validateManifest()` 並傳入 `{ weight: NaN }` 應該傳回新的物件 `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }`，且不會修改原始輸入。你可以使用 `Number.isNaN()` 來驗證 NaN 值。

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated");
```

如果輸入的清單物件無效，你的 `validateManifest` 函式應該傳回描述遺漏和／或無效屬性的物件。

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

你的 `validateManifest` 函式應該傳回新的物件，且不會修改原始物件。

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

你應該有一個名為 `processManifest` 的函式，帶有一個 `manifest` 參數。

```js
assert.isFunction(processManifest);
const funcString = processManifest.toString();
const params = __helpers.getFunctionParams(funcString);
assert.equal(params.length, 1);
assert.equal(params[0].name, 'manifest');
```

呼叫 `processManifest()` 並傳入 `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }` 應該先記錄 `"Validation success: 55"`，然後再記錄 `"Total weight: 180 kg"`。

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

如果輸入的清單物件有效，你的 `processManifest` 函式應該先記錄成功訊息 `Validation success: ${containerId}`。

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

如果輸入的清單物件有效，你的 `processManifest` 函式應該使用 `normalizeUnits()` 將其標準化為公斤，然後記錄：`Total weight: ${weight} kg`。

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ]
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

如果輸入的清單物件有效，你的 `processManifest` 函式應該記錄帶有該物件 `containerId` 的成功訊息，然後記錄該物件的 `weight`（以公斤為單位）。你應該使用 `normalizeUnits()` 進行轉換，並有兩次 `console.log()` 呼叫。

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call normalizeUnits() when the manifest is valid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ]
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

呼叫 `processManifest()` 並傳入 `{ containerId: -88, destination: "Soledad", weight: NaN }` 應該先記錄 `Validation error: -88`，然後記錄物件 `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }`。

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

呼叫 `processManifest()` 並傳入 `{ destination: "Watsonville", hazmat: true }` 應該先記錄 `Validation error: undefined`，然後記錄物件 `{ containerId: "Missing", weight: "Missing", unit: "Missing" }`。

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

如果輸入的清單物件無效，你的 `processManifest` 函式應該先記錄錯誤訊息 `Validation error: ${containerId}`。

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

如果輸入的清單物件無效，你的 `processManifest` 函式也應該記錄呼叫 `validateManifest()` 並傳入原始清單物件所返回的物件。直接呼叫 `console.log()` 並傳入該物件。

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

如果輸入的清單物件無效，你的 `processManifest` 函式應該記錄帶有該物件 `containerId` 的錯誤訊息，然後記錄呼叫 `validateManifest()` 並傳入輸入物件所返回的物件。這應該包含兩次 `console.log()` 呼叫。

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest should call validateManifest() when the manifest is invalid"
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);   
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
