---
id: 66a9b333487b9c14998539a5
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside your `head` element, nest a `meta` element with the `charset` attribute set to the value `"UTF-8"`.

Додайте елемент `title` під цим елементом `meta`.

Текстом елемента `title` має бути `Calculus Final Exams Table`.

After you complete your `head` element, you can add your `body` element.

# --hints--

You should have a `meta` element.

```js
assert.isNotNull(document.querySelector("meta"));
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Атрибут `charset` повинен мати значення `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Your `meta` element should be nested inside your `head` element.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Ви повинні мати початковий тег `title`.

```js
assert.match(code, /<title>/i);
```

Ви повинні мати кінцевий тег `title`.

```js
assert.match(code, /<\/title>/i);
```

Your `title` element should be nested in your `head` element.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Елемент `title` повинен містити текст `Calculus Final Exams Table`. You may need to check your spelling.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

Ви повинні мати початковий тег `body`.

```js
assert.match(code, /<body>/i);
```

Ви повинні мати кінцевий тег `body`.

```js
assert.match(code, /<\/body>/i);
```

Не змінюйте елемент `head`. Make sure you did not delete your closing tag.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Your `body` element should come after your `head` element.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
