---
id: 66a9bb8578459a175432e7d0
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

For the first section of the table, you will want to group the header content which represents the column labels for the student's first name, last name, and final exam grade.

The `table head` element, `<thead>`, is used to group the header content in a table.

Ось приклад використання елемента `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Додайте елемент `table head` під елементом `caption`.

# --hints--

Ви повинні мати початковий тег `thead`.

```js
assert.match(code, /<thead>/i);
```

Ви повинні мати кінцевий тег `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Елемент `thead` має бути в межах елемента `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl.parentElement.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
      --fcc-editable-region--

      --fcc-editable-region--
    </table>
  </body>
</html>
```
