---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

Настав час додати ще двох студентів до таблиці.

Дотримуючись шаблону з попереднього кроку, додайте рядок для другого студента. Використайте такі дані для таблиці:

- Прізвище: `Doe`
- Ім’я: `Samantha`
- Оцінка: `92`

Додайте ще один рядок з третім студентом, використовуючи такі дані:

- Прізвище: `Rodriguez`
- Ім’я: `Marcus`
- Оцінка: `88`

# --hints--

В межах елемента `tbody` має бути другий елемент `tr`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Ви повинні мати елемент `td` з текстом `Doe` у другому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent, 'Doe');
```

Ви повинні мати елемент `td` з текстом `Samantha` у другому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent, 'Samantha');
```

Ви повинні мати елемент `td` з текстом `92` у другому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent, '92');
```

В межах елемента `tbody` має бути третій елемент `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Ви повинні мати елемент `td` з текстом `Rodriguez` у третьому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent, 'Rodriguez');
```

Ви повинні мати елемент `td` з текстом `Marcus` у третьому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent, 'Marcus');
```

Ви повинні мати елемент `td` з текстом `88` у третьому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent, '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
