---
id: 66a9ca099e3ecb1be10a2696
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Now it is time to add two more students to the table.

Following the same pattern as the previous steps, add a fourth student table row. Use the following data for the table data elements:

- Last Name: `Thompson`
- First Name: `Jane`
- Grade: `77`

For the fifth student table row, use the following data for the table data elements:

- Last Name: `Williams`
- First Name: `Natalie`
- Grade: `83`

# --hints--

You should have a fourth `tr` inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Ви повинні мати елемент `td` з текстом `Thompson` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent, 'Thompson');
```

Ви повинні мати елемент `td` з текстом `Jane` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent, 'Jane');
```

Ви повинні мати елемент `td` з текстом `77` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent, '77');
```

You should have a fifth `tr` inside your `tbody` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Ви повинні мати елемент `td` з текстом `Williams` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent, 'Williams');
```

Ви повинні мати елемент `td` з текстом `Natalie` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent, 'Natalie');
```

Ви повинні мати елемент `td` з текстом `83` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent, '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
