---
id: 66a9b48a211a73155621d0df
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

In the previous lectures, you learned how to work with the `table` element to represent tabular data.

Inside your `body` element, nest a `table` element.

# --hints--

Ви повинні мати початковий тег `table`.

```js
assert.match(code, /<table>/i);
```

Ви повинні мати кінцевий тег `table`.

```js
assert.match(code, /<\/table>/i);
```

Your `table` element should be inside your `body` element.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```
