---
id: 66a9bcf00f13a418368a272e
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Заголовок таблиці складається з елемента `tr` (скорочено від <dfn>table row</dfn>), який містить елементи `th` (тобто <dfn>комірки заголовка таблиці</dfn>).

Ось приклад використання елементів `tr` та `th` в спортивній таблиці:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Додайте елемент `tr` до елемента `thead`.

Додайте три елементи `th` до елемента `tr`.

Перший елемент `th` повинен мати текст `Last Name`. Другий елемент `th` повинен мати текст `First Name`. Третій елемент `th` повинен мати текст `Grade`.

# --hints--

Ви повинні мати початковий тег `tr`.

```js
assert.match(code, /<tr>/i);
```

Ви повинні мати кінцевий тег `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Елемент `tr` має бути в межах елемента `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Перший елемент `th` повинен мати текст `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent, 'Last Name');
```

Другий елемент `th` повинен мати текст `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent, 'First Name');
```

Третій елемент `th` повинен мати текст `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent, 'Grade');
```

Усі елементи `th` мають бути в межах елемента `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      --fcc-editable-region--
      <thead>     

      </thead>
      --fcc-editable-region--
    </table>
  </body>
</html>
```
