---
id: 66a9cb26ec6bd41cf6c82bc5
title: Крок 11
challengeType: 0
dashedName: step-11
---

# --description--

Останнім розділом буде підсумкова частина таблиці — елемент `tfoot` (скорочено від <dfn>table foot</dfn>). Цей елемент буде використано для середньої оцінки всіх студентів у таблиці.

Додайте елемент `tfoot` до таблиці.

Додайте елемент `tr` до елемента `tfoot`.

Додайте два елементи `td` до елемента `tr`.

Перший елемент `td` повинен мати текст `Average Grade`. Другий елемент `td` повинен мати текст `78.8`.

# --hints--

Ви повинні мати початковий тег `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Ви повинні мати кінцевий тег `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Ви повинні мати елемент `tr` в межах елемента `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Перший елемент `td` повинен мати текст `Average Grade`. Переконайтеся, що він знаходиться в елементі `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent, 'Average Grade');
```

Другий елемент `td` повинен мати текст `78.8`. Переконайтеся, що він знаходиться в елементі `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent, '78.8');
```

Елемент `tr` повинен мати два елементи `td`.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
