---
id: 66a9b48a211a73155621d0df
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У попередніх лекціях ви навчилися працювати з елементом `table`, щоб представляти табличні дані.

Додайте елемент `table` до елемента `body`.

# --hints--

Ви повинні мати початковий тег `table`.

```js
assert.match(code, /<table>/i);
```

Ви повинні мати кінцевий тег `table`.

```js
assert.match(code, /<\/table>/i);
```

Елемент `table` має бути в межах елемента `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
