---
id: 66a9b8f14b963916a3baa732
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

To add a caption to a `table`, you can use the <dfn>table caption</dfn> element.

Here is an example using the `caption` element:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Inside your `table` element, nest a `caption` element with the text `Calculus Final Exam Grades`.

# --hints--

Ви повинні мати початковий тег `caption`.

```js
assert.match(code, /<caption>/i);
```

Ви повинні мати кінцевий тег `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Елемент `caption` має бути в межах елемента `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

Your `caption` should have the text `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText, 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    --fcc-editable-region--
    <table>

    </table>
    --fcc-editable-region--
  </body>
</html>
```
