---
id: 66a9c14e3b34c719e34bae20
title: Крок 8
challengeType: 0
dashedName: step-8
---

# --description--

Щоб додати дані студента до таблиці, потрібно використати елементи `tr` (скорочено від <dfn>table row</dfn>) та `td` (скорочено від <dfn>table data</dfn>).

Для створення комірки в таблиці використовують елемент `td`.

Ось приклад використання елемента `td` для таблиці з гравцями:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

До тіла таблиці додайте елемент `tr`, який містить три елементи `td`.

Перший елемент `td` повинен містити прізвище `Davis`.

Другий елемент `td` повинен містити ім’я `Alex`.

Третій елемент `td` повинен містити оцінку `54`.

# --hints--

В межах елемента `tbody` має бути елемент `tr`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Ви повинні мати елемент `td` з текстом `Davis` в елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent, 'Davis');
```

Ви повинні мати елемент `td` з текстом `Alex` в елементі `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent, 'Alex');
```

Ви повинні мати елемент `td` з текстом `54` в елементі `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent, '54');
```

Усі три елементи `td` мають бути в межах елемента `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
