---
id: 66a9b333487b9c14998539a5
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

В межах елемента `head` вкладіть елемент `meta` з атрибутом `charset` зі значенням `"UTF-8"`.

Додайте елемент `title` під цим елементом `meta`.

Текстом елемента `title` має бути `Calculus Final Exams Table`.

Після завершення елемента `head` можна додати елемент `body`.

# --hints--

Ви повинні мати елемент `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

Елемент `meta` є порожнім елементом, він не повинен мати кінцевий тег `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Тег `meta` повинен мати атрибут `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Атрибут `charset` повинен мати значення `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

Елемент `meta` має бути вкладеним в елементі `head`.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Ви повинні мати початковий тег `title`.

```js
assert.match(code, /<title>/i);
```

Ви повинні мати кінцевий тег `title`.

```js
assert.match(code, /<\/title>/i);
```

Елемент `title` має бути вкладеним в елементі `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

Елемент `title` повинен мати текст `Calculus Final Exams Table`. Можливо, доведеться перевірити написання.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Calculus Final Exams Table");
```

Ви повинні мати початковий тег `body`.

```js
assert.match(code, /<body>/i);
```

Ви повинні мати кінцевий тег `body`.

```js
assert.match(code, /<\/body>/i);
```

Не змінюйте елемент `head`. Переконайтеся, що не видалили кінцевий тег.

```js
assert.match(code, /<head>/i);
assert.match(code, /<\/head>/i);
```

Елемент `body` має бути після елемента `head`.

```js
assert.match(code, /<\/head>[.\n\s]*<body>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
