---
id: 66a9c006e4f58819396605b3
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Ви завершили заголовок, тому можете додати тіло таблиці — елемент `tbody` (скорочено від <dfn>table body</dfn>). Тіло таблиці представлятиме імена студентів та їхні оцінки.

Додайте елемент `tbody` до таблиці.

# --hints--

Ви повинні мати початковий тег `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Ви повинні мати кінцевий тег `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Початковий тег `tbody` має бути перед кінцевим тегом `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
