---
id: 66a9ca099e3ecb1be10a2696
title: Крок 10
challengeType: 0
dashedName: step-10
---

# --description--

Настав час додати ще двох студентів до таблиці.

Дотримуючись шаблону з попередніх кроків, додайте рядок для четвертого студента. Використайте такі дані для таблиці:

- Прізвище: `Thompson`
- Ім’я: `Jane`
- Оцінка: `77`

Додайте ще один рядок з п’ятим студентом, використовуючи такі дані:

- Прізвище: `Williams`
- Ім’я: `Natalie`
- Оцінка: `83`

# --hints--

В межах елемента `tbody` має бути четвертий елемент `tr`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Ви повинні мати елемент `td` з текстом `Thompson` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent, 'Thompson');
```

Ви повинні мати елемент `td` з текстом `Jane` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent, 'Jane');
```

Ви повинні мати елемент `td` з текстом `77` у четвертому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent, '77');
```

В межах елемента `tbody` має бути п’ятий елемент `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Ви повинні мати елемент `td` з текстом `Williams` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent, 'Williams');
```

Ви повинні мати елемент `td` з текстом `Natalie` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent, 'Natalie');
```

Ви повинні мати елемент `td` з текстом `83` у п’ятому елементі `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent, '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
