---
id: 6551f4259d66026ee28cb205
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

يمكنك أيضًا الوصول إلى حروف السلسلة النصية بدءًا من نهاية السلسلة. الحرف الأخير له فهرس `-1`، والحرف الذي يسبقه `-2` وهكذا.

عدّل استدعاء `print()` الحالي ليطبع الحرف الأخير في السلسلة النصية.

# --hints--

يجب أن تستدعي دالة `print()` كما في السابق.

```js
assert.match(code, /print\s*\(.*\)/)
```

مرر `text[-1]` إلى دالة `print()` بوضعه بين الأقواس العادية. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
