---
id: 6551f5756c20146fc94f8675
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

يمكنك الوصول إلى عدد الأحرف في سلسلة نصية باستخدام الدالة المدمجة `len()`.

عدّل استدعاء `print()` الحالي بتمرير `len(text)` بدلًا من `text[-1]`.

# --hints--


يجب استدعاء الدالة `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

يجب تمرير المتغير `text` إلى الدالة `len()` بوضعه بين الأقواس العادية.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

يجب تمرير `len(text)` إلى الدالة `print()` بوضعها بين الأقواس العادية. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

يجب ألا يحتوي شفرتك على `print(text[-1])`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
