---
id: 6551fe3b1df7c9740f13f270
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

الجوانب الأساسية لتسمية المتغيرات في Python هي:

- بعض الكلمات هي كلمات مفتاحية محجوزة (مثل `for`، `while`، `True`). لها معنى خاص في Python، لذا لا يمكنك استخدامها كأسماء متغيرات.
- لا يمكن أن تبدأ أسماء المتغيرات برقم، ويجب أن تحتوي فقط على أحرف أبجدية رقمية أو شرطات سفلية.
- أسماء المتغيرات حساسة لحالة الأحرف، أي أن `my_var` يختلف عن `my_Var` و`MY_VAR`.
- أخيرًا، من الاصطلاحات الشائعة كتابة أسماء المتغيرات باستخدام `snake_case`، حيث يُستبدل كل فراغ بشرطة سفلية وتُكتب الكلمات بأحرف صغيرة.

احذف كلا الاستدعاءين لـ `print()` وأعلن متغيرًا آخر اسمه `alphabet`. عيّن السلسلة النصية `'abcdefghijklmnopqrstuvwxyz'` لهذا المتغير.

# --hints--

يجب ألا يحتوي شفرتك على `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

يجب ألا يحتوي شفرتك على `print(type(shift))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

يجب إعلان متغير اسمه `alphabet`. انتبه لوضع اسم المتغير في بداية السطر.

```js
assert.match(code, /^alphabet\s*=/m)
```

يجب تعيين السلسلة النصية `'abcdefghijklmnopqrstuvwxyz'` لمتغير `alphabet`. تذكر استخدام علامات اقتباس مفردة أو علامات اقتباس مزدوجة لإحاطة السلسلة النصية.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
