---
id: 6551f3bfed58796ebb1268dc
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

يمكن الإشارة إلى كل حرف في السلسلة النصية بواسطة فهرس رقمي. يبدأ عد الفهارس من الصفر. لذلك، الحرف الأول في السلسلة النصية له فهرس `0`. على سبيل المثال، في السلسلة النصية `'Hello World'`، الحرف `'H'` في الفهرس `0`، والحرف `'e'` في الفهرس `1`، وهكذا.

يمكن الوصول إلى كل حرف في السلسلة النصية باستخدام ترميز الأقواس. يجب كتابة اسم المتغير متبوعًا بأقواس مربعة، مع إضافة فهرس الحرف بين الأقواس:

```py
text = 'Hello World'
r = text[8]
```

الآن، بدلًا من طباعة `text`، اطبع فقط الحرف الموجود في الفهرس `6`.

# --hints--

يجب أن تستدعي دالة `print()` كما هو.

```js
assert.match(code, /print\s*\(.*\)/)
```

يجب تمرير `text[6]` إلى دالة `print()` بوضعه بين الأقواس العادية. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

تحتوي الشفرة الخاصة بك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
