---
id: 6551f6f9ec6db270eb83f3db
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

يمكنك رؤية `11` مطبوعة على الطرفية لأن `'Hello World'` يحتوي على 11 حرفًا.

دالة مدمجة مفيدة أخرى هي `type()`، التي تُرجع نوع البيانات لمتغير. عدل استدعاء `print()` ليطبع نوع البيانات لـ `text`.

# --hints--

يجب استدعاء الدالة `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

يجب تمرير `text` إلى الدالة `type()` بوضعه بين الأقواس العادية.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

يجب تمرير `type(text)` إلى الدالة `print()` بوضعه بين الأقواس العادية. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

يجب ألا يحتوي شفرتك على `print(len(text))`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
