---
id: 6551ecf597038c36637f3759
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

المتغيرات جزء أساسي من Python وأي لغة برمجة. المتغير هو اسم يشير إلى كائن أو يمثل مرجعًا له. يمكنك إعلان متغير بكتابة اسم المتغير على الجانب الأيسر من معامل التعيين `=` وتحديد القيمة التي تريد تعيينها لذلك المتغير على الجانب الأيمن من معامل التعيين:

```py
variable_name = value
```

أنشئ متغيرًا اسمه `number` وعيّن القيمة `5` إلى متغيرك الجديد.

# --hints--

يجب إعلان متغير اسمه `number`. انتبه لوضع اسم المتغير في بداية السطر.

```js
assert.match(code, /^number\s*=/m)
```

يجب تعيين القيمة `5` إلى متغير `number` الخاص بك.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
