---
id: 6551f1e8737dbe6e0cb08ea4
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

والآن اطبع المتغير الجديد.

# --hints--

لا يجب إزالة استدعاء `print(type(text))` الحالي.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

يجب استدعاء الدالة `print()` مع تمرير المتغير `shift`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

تحتوي الشفرة على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
