---
id: 6551f22653e6ce6e4f1d1370
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

عدّل استدعاء `print(shift)` ليطبع نوع البيانات للمتغير `shift`.

# --hints--

يجب أن تحتفظ باستدعاء `print(type(text))` الحالي.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

يجب أن تمرر `shift` إلى الدالة `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

يجب أن تمرر `type(shift)` إلى دالة `print()`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

يجب ألا يحتوي شفرتك على `print(shift)`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
