---
id: 6551ef504d91936d2d4e54f8
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

المعلمة هي كائن أو تعبير يُمرر إلى دالة — تُضاف بين الأقواس العادية المفتوحة والمغلقة — عند استدعائها:

```py
greet = 'Hello!'
print(greet)
```

الشفرة في المثال أعلاه ستطبع السلسلة النصية `'Hello!'`، وهي قيمة المتغير `greet` الذي يُمرر إلى دالة `print()` كمعلمة.

اطبع متغير `text` على الشاشة بتمرير متغير `text` كمعلمة إلى دالة `print()`.

# --hints--

يجب تمرير `text` إلى دالة `print()` بإضافة اسم هذا المتغير داخل الأقواس العادية المفتوحة والمغلقة. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

تحتوي الشفرة الخاصة بك على صيغة نحوية غير صحيحة و/أو تنسيق غير صحيح للمسافات البادئة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
