---
id: 6551fb531ce7f773071a0b84
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

كما ترى، ناتج طباعة `type(text)` هو `<class 'str'>`، مما يعني أن متغيرك هو سلسلة نصية، موضحًا بـ `str`.

انتقل الآن إلى سطر جديد وأنشئ متغيرًا آخر اسمه `shift` وقم بتعيين القيمة `3` لهذا المتغير.

# --hints--

يجب إعلان متغير اسمه `shift`. انتبه لوضع اسم المتغير في بداية السطر.

```js
assert.match(code, /^shift\s*=/m)
```

يجب تعيين القيمة `3` لمتغير `shift` الخاص بك.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
