---
id: 655261b2e1f2c197093f3993
title: الخطوة 41
challengeType: 20
dashedName: step-41
---

# --description--

حاليًا، تُشفّر الفراغات كـ `'c'`. للحفاظ على الفراغات الأصلية في الرسالة النصية، تحتاج إلى جملة شرطية `if`. تتكون هذه الجملة من الكلمة المفتاحية `if`، شرط، ونقطتين رأسيتين `:`.

```py
if x != 0:
    print(x)
```

في المثال أعلاه، شرط جملة الـ `if` هو `x != 0`. الشفرة `print(x)` داخل جسم جملة الـ `if` تُنفذ فقط عندما يُقيّم الشرط إلى `True` (في هذا المثال، يعني أن `x` مختلف عن الصفر).

في أعلى حلقة for، استبدل `print(char == ' ')` بجملة `if`. يجب أن تتحقق شرط جملة `if` ليكون `True` إذا كانت قيمة `char` فراغًا، و`False` في غير ذلك. داخل جسم جملة `if`، اطبع السلسلة النصية `'space!'`. تذكر أن تضع مسافة بادئة لهذا السطر.

# --hints--

يجب ألا يحتوي شفرتك على `print(char == ' ')`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

يجب أن تستبدل `print(char == ' ')` بجملة `if` تُنفذ عندما يكون `char == ' '`. لا تستخدم أقواسًا عادية حول شرط `if` وتذكر تضمين النقطتين الرأسيتين في النهاية.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

يجب أن تطبع السلسلة النصية `'space!'` داخل جملة `if` الجديدة.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
