---
id: 6552487e689f6e895f658717
title: خطوة 33
challengeType: 20
dashedName: step-33
---

# --description--

الآن عليك إنشاء متغير `new_char` في نهاية جسم الحلقة. عيّن قيمته إلى `alphabet[new_index]`.

# --hints--

يجب إنشاء متغير `new_char` داخل حلقة `for` الخاصة بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

يجب تعيين متغير `new_char` إلى `alphabet[new_index]` في نهاية جسم الحلقة.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
