---
id: 6552008c0d9d9075cbec9772
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

ستستخدم طريقة `.find()` لإيجاد موقع كل حرف في الأبجدية داخل رسالتك. الطريقة تشبه الدالة، لكنها تنتمي إلى كائن.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

في الأعلى، تُستدعى طريقة `.find()` على `sentence` (السلسلة النصية التي يُبحث فيها)، ويتم تمرير `'r'` (الحرف المراد إيجاده) كمعلمة. استدعاء `sentence.find('r')` سيُرجع `4`، وهو فهرس أول ظهور للحرف `'r'` في `sentence`.

في نهاية شفرتك، استدعِ `.find()` على `alphabet` ومرر `'z'` كمعلمة للطريقة.

# --hints--

يجب أن تستدعي طريقة `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

يجب أن تستدعي طريقة `.find()` على المتغير `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


يجب أن تستدعي طريقة `.find()` على المتغير `alphabet` وتمرر `'z'` للطريقة. انتبه لوضع استدعاء الطريقة في بداية السطر.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تنسيق مسافات بادئة غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
