---
id: 65522255d5b9cd80f335c6f2
title: الخطوة 29
challengeType: 20
dashedName: step-29
---

# --description--

في نهاية جسم الحلقة، أعلن عن متغير باسم `new_index` وقم بتعيين قيمة `index + shift` لهذا المتغير.

# --hints--

يجب إعلان متغير باسم `new_index` داخل حلقة `for` الخاصة بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

يجب تعيين `index + shift` للمتغير الجديد في نهاية جسم حلقة `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
