---
id: 65520e6f2b9678799977f24d
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

احذف آخر استدعاء لـ `print()`. ثم، بدلًا من `text[0]`، مرر `text[0].lower()` كمعلمة لاستدعاء `.find()` الخاص بك واطلع على النتيجة.

# --hints--

يجب ألا يحتوي شفرتك على `print(text.lower())`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

يجب أن تطبع متغير `index` الخاص بك أيضًا.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

يجب تحديث استدعاء `alphabet.find(text[0])` ليستخدم `text[0].lower()` كمعلمة. انتبه لوضع استدعاء الطريقة في بداية السطر.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
