---
id: 655209a4a27dd37873c4cac3
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

تُعيد `.find()` فهرس الحرف المطابق داخل السلسلة النصية. إذا لم يُعثر على الحرف، تُعيد `-1`. كما ترى، الحرف الأول في `text`، وهو `'H'` كبير، غير موجود، لأن `alphabet` يحتوي فقط على أحرف صغيرة.

يمكنك تحويل السلسلة النصية إلى ما يعادلها من أحرف صغيرة باستخدام الطريقة `.lower()`. أضف استدعاء `print()` آخر لطباعة `text.lower()` ومشاهدة النتيجة.

# --hints--

يجب أن يبقى لديك `print(index)` في شفرتك. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

يجب أن تطبع `text.lower()`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تنسيق مسافات بادئة غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
