---
id: 655251308f31958d06cdf267
title: الخطوة 36
challengeType: 20
dashedName: step-36
---

# --description--

في الوقت الحالي، يتم تحديث الحرف المشفر في كل تكرار. من الأفضل تخزين السلسلة النصية المشفرة في متغير جديد. قبل حلقة `for`، أعلن عن متغير يسمى `encrypted_text` وقم بتعيين سلسلة نصية فارغة (`''`) لهذا المتغير.

# --hints--

يجب إعلان متغير يسمى `encrypted_text` قبل حلقة `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

يجب تعيين سلسلة نصية فارغة للمتغير `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
