---
id: 65520f8282faf57a0db4f7fe
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

كما ترى من الناتج، الحرف `'h'` موجود في الفهرس `7` في سلسلة النصية `alphabet`. الآن تحتاج إلى إيجاد الحرف عند الفهرس `7` مضافًا إليه قيمة `shift`. لهذا، يمكنك استخدام عامل الجمع `+` بنفس الطريقة التي تستخدمها في الجمع الرياضي.

عدّل المتغير `shifted` ليخزن قيمة `alphabet` عند الفهرس `index + shift`.

# --hints--

يجب تعيين `alphabet[index + shift]` إلى متغير `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
