---
id: 655220953ba90d80514d7ee2
title: الخطوة 27
challengeType: 20
dashedName: step-27
---

# --description--

حاليًا، دالة `print()` تأخذ معلمة واحدة فقط وهي `char`، لكنها يمكن أن تأخذ عدة معلمات مفصولة بفاصلة.

أضف معلمة ثانية إلى `print(char)` بحيث تطبع الحرف وفهرسه داخل الأبجدية.

# --hints--

يجب أن تضيف `index` كمعلمة ثانية إلى استدعاء `print(char)` الحالي. لا تنسَ فصل المعلمات بفاصلة.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
