---
id: 655208d59b131e7816f18c96
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

تُعطي الدالة `print()` مخرجات فقط في وحدة التحكم، لكن الدوال والطرق يمكن أن تحتوي على قيمة إرجاع يمكنك استخدامها في شفرتك.

الآن عيّن `alphabet.find(text[0])` إلى متغير اسمه `index`. بهذه الطريقة، سيخزن `index` القيمة التي تُرجعها `alphabet.find(text[0])`.

# --hints--

يجب إعلان متغير اسمه `index`. انتبه لوضع اسم المتغير في بداية السطر.

```js
assert.match(code, /^index\s*=/m)
```

يجب تعيين القيمة التي تُرجعها `alphabet.find(text[0])` إلى متغير `index` الخاص بك.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
