---
id: 65521fc818947e800bffe48a
title: الخطوة 26
challengeType: 20
dashedName: step-26
---

# --description--

داخل حلقة `for`، قبل طباعة الحرف الحالي، أعلن عن متغير اسمه `index` وقم بتعيين القيمة التي تُرجعها `alphabet.find(char)` لهذا المتغير.

# --hints--

يجب إعلان متغير جديد اسمه `index` في بداية حلقة `for` الخاصة بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

يجب تعيين `alphabet.find(char)` إلى متغير `index` الجديد الخاص بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
