---
id: 65521203d7165c7b84b22ad4
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

تسمح الحلقة بالتكرار على متتالية من العناصر وتنفيذ إجراءات على كل عنصر منها بشكل منهجي.

في هذه الحالة، ستستخدم حلقة `for`. إليك كيف يمكنك التكرار على `text`:

```py
for i in text:
```

`for` هي الكلمة المفتاحية التي تحدد نوع الحلقة. `i` متغير يأخذ بالتتابع قيمة عناصر `text`. تنتهي العبارة بنقطتين رأسيتين، `:`.

أسفل السطر الذي أعلنت فيه `alphabet`، اكتب حلقة `for` للتكرار على `text`. استخدم `i` كمتغير للحلقة.

عند القيام بذلك، يظهر خطأ في الطرفية. ستتعلم عنه في الخطوة التالية.

# --hints--

يجب استخدام الكلمة المفتاحية `for` لإنشاء حلقة. تأكد من وضع `for` في بداية السطر وترك مسافة بيضاء بعدها.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

يجب كتابة المتغير `i` بعد الكلمة المفتاحية `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

يجب كتابة الكلمة المفتاحية `in` بعد `for i `. تأكد من ترك مسافة حول كلمة `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

يجب كتابة `text` بعد `for i in `. لا تنسَ إضافة النقطتين الرأسيتين النهائيتين `:`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

يجب وضع حلقة `for` أسفل السطر الذي يحتوي على الشفرة `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
