---
id: 65525e359ca28d938baa82c5
title: الخطوة 39
challengeType: 20
dashedName: step-39
---

# --description--

يمكنك الحصول على نفس تأثير `a = a + b` باستخدام معامل التعيين بالإضافة:

```py
a += b
```

يمكنك معامل التعيين بالإضافة من إضافة قيمة إلى متغير ثم تعيين النتيجة لذلك المتغير.

استخدم المعامل `+=` لإضافة قيمة وتعيينها في نفس الوقت إلى `encrypted_text`.

# --hints--

يجب استخدام معامل التعيين بالإضافة لإضافة `alphabet[new_index]` إلى القيمة الحالية لـ `encrypted_text` داخل جسم الحلقة.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
