---
id: 65524d2a1a253b8bb5197ae2
title: الخطوة 28
challengeType: 20
dashedName: step-28
---

# --description--

تُعيد الدالة `find` مرة أخرى القيمة `-1` للحروف الكبيرة، وكذلك لحرف المسافة. ستتعامل مع المسافة لاحقًا.

في الوقت الحالي، بدلًا من التكرار على `text`، غيّر حلقة `for` لتتكرر على `text.lower()`.

# --hints--

يجب تعديل حلقة `for` لتتكرر على `text.lower()` بدلًا من `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
