---
id: 6553f6b4ec51112d44d737c8
title: الخطوة 48
challengeType: 20
dashedName: step-48
---

# --description--

قبل استدعاء `print` مباشرة، أضف استدعاءً آخر ومرر `'plain text:', text` كالمعلمات إلى `print()`. استخدم نفس المسافة البادئة.

# --hints--

يجب أن تطبع `'plain text:', text` بعد حلقة for الخاصة بك.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
