---
id: 6553ed69ece88d29594748aa
title: الخطوة 52
challengeType: 20
dashedName: step-52
---

# --description--

لتنفيذ دالة، يجب استدعاؤها (أو تنفيذها) بإضافة زوج من الأقواس العادية بعد اسمها، مثل هذا:

```py
function_name()
```

في نهاية شفرتك، استدعِ دالة `caesar`. انتبه إلى المسافة البادئة.

# --hints--

يجب أن تستدعي دالة `caesar`. تأكد من كتابة استدعاء الدالة في بداية السطر.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
