---
id: 6553a755879b131a445e664c
title: الخطوة 44
challengeType: 20
dashedName: step-44
---

# --description--

حاول تعيين السلسلة النصية `'Hello Zaira'` إلى المتغير `text` وشاهد ما يحدث في الطرفية.

سترى استثناء `string index out of range`. لا تقلق، ستعرف كيف تصلحه قريبًا!

# --hints--

يجب أن يكون لديك متغير `text`.

```js
assert.match(code, /^text\s*=/m)
```

يجب تعيين السلسلة النصية `'Hello Zaira'` إلى المتغير `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
