---
id: 6553f4f66099802c6ae94613
title: الخطوة 55
challengeType: 20
dashedName: step-55
---

# --description--

حاليًا، شفرتك ترفع خطأ `TypeError`، لأن دالة `caesar` معرفة بمعلمتين (`message` و `offset`)، لذلك تتوقع أن تُستدعى بمعلمتين.

استدعاء `caesar()` بدون المعلمات المطلوبة يوقف تنفيذ الشفرة.

أعطِ قيمًا لـ `message` و `offset` بتمرير `text` و `shift` كمعلمات لاستدعاء دالة `caesar`.

# --hints--

يجب تمرير `text` و `shift` كمعلمات لاستدعاء دالتك بوضعهما داخل الأقواس العادية. لا تنسَ فصل المعلمات بفاصلة.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
