---
id: 6553f6086add4b2cbb99fd78
title: الخطوة 56
challengeType: 20
dashedName: step-56
---

# --description--

في أسفل شفرتك، بعد استدعاء `caesar(text, shift)` الحالي، استدعِ دالتك مرة أخرى.

هذه المرة، مرر المتغير `text` والعدد الصحيح `13` كمعلمات.

# --hints--

يجب أن تستدعي دالتك مرة أخرى، هذه المرة مع تمرير `text` و `13` كمعلمات.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

يجب أن تحتفظ باستدعاء الدالة الحالي.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
