---
id: 6553efd6ada3f42aa2d75448
title: الخطوة 53
challengeType: 20
dashedName: step-53
---

# --description--

الآن يجب أن ترى الناتج مرة أخرى. رغم أن هذه الطريقة تعمل، إلا أنها لا تعزز بشكل كبير إعادة استخدام الشفرة. استدعاء دالتك مرارًا سيؤدي إلى نفس النتيجة. مع ذلك، يمكن إعلان الدوال مع *معلمات* لإضافة المرونة والتخصيص:

```py
def function_name(param_1, param_2):
    <code>
```

المعلمات هي متغيرات يمكنك استخدامها داخل دالتك. يمكن إعلان دالة بعدد مختلف من المعلمات. في المثال أعلاه، `param_1` و `param_2` هما معلمتان.

عدّل إعلان دالتك بحيث تأخذ معلمتين تسمى `message` و `offset`.

بعد ذلك، سترى خطأ يظهر في الطرفية. سترى كيفية حله في الخطوات التالية.

# --hints--

يجب أن تأخذ دالة `caesar` المعلمتين `message` و `offset`. تذكر فصل المعلمات بفاصلة.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
