---
id: 6553995f412dd8122ed38e4a
title: الخطوة 43
challengeType: 20
dashedName: step-43
---

# --description--

يمكن أن يحتوي البيان الشرطي أيضًا على بند `else`. يمكن إضافة هذا البند إلى نهاية بيان `if` لتنفيذ شفرة بديلة إذا كانت شرط بيان `if` خاطئًا:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

كما ترى في الناتج، عندما تصل تكرارات الحلقة إلى الفراغ، يُضاف فراغ إلى السلسلة النصية المشفرة. ثم تُنفذ الشفرة خارج كتلة `if` ويُضاف الحرف `c` إلى السلسلة النصية المشفرة.

لإصلاح ذلك، أضف بند `else` بعد `encrypted_text += char` وابدأ بمسافة بادئة لجميع الأسطر التالية من الشفرة ما عدا استدعاء `print()`.

# --hints--

يجب أن تنشئ بند `else`. تذكر تضمين النقطتين الرأسيتين في النهاية.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

يجب أن تبدأ الأسطر التي تلي بند `else` بمسافة بادئة ما عدا استدعاء `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

تحتوي شفراتك على صيغة نحوية غير صحيحة و/أو مسافات بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
