---
id: 6553980e0527fa115c705646
title: الخطوة 42
challengeType: 20
dashedName: step-42
---

# --description--

الآن، بدلًا من طباعة `'space!'`، استخدم معامل التعيين بالإضافة لإضافة المسافة (المخزنة حاليًا في `char`) إلى القيمة الحالية لـ `encrypted_text`.

# --hints--

يجب ألا يحتوي شفرتك على `print('space!')`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

يجب استخدام معامل `+=` لإضافة `char` إلى القيمة الحالية لـ `encrypted_text` داخل جملة `if` الجديدة.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
