---
id: 6553a572f7a65718f1e42e18
title: الخطوة 49
challengeType: 20
dashedName: step-49
---

# --description--

الدالة هي في الأساس كتلة شفرة قابلة لإعادة الاستخدام. لقد تعرفت بالفعل على بعض الدوال المدمجة، مثل `print()` و `find()` و `len()`. لكن يمكنك أيضًا تعريف دوال مخصصة مثل هذه:

```py
def function_name():
    <code>
```

يبدأ إعلان الدالة بالكلمة المفتاحية `def` متبوعة باسم الدالة — وهو اسم متغير صالح — وزوج من الأقواس العادية. ينتهي الإعلان بنقطتين رأسيتين.

بعد متغير `shift` مباشرة، أعلن دالة باسم `caesar` وادخل جميع الأسطر التالية لتعطي دالتك الجديدة جسمًا.

# --hints--

يجب أن تستخدم الكلمة المفتاحية `def` لإعلان دالة جديدة.

```js
assert.match(code, /^def\s+/m)
```

يجب أن تكتب `caesar` كاسم الدالة بعد الكلمة المفتاحية `def`. تذكر إضافة مسافة بعد `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

يجب أن تضيف زوجًا من الأقواس العادية بعد اسم الدالة. لا تنسَ النقطتين الرأسيتين في النهاية.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

يجب أن تدخل جميع الأسطر بعد `shift = 3` لتصبح جسم دالتك الجديدة.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
