---
id: 6553f8c570f9982e013a8886
title: الخطوة 57
challengeType: 20
dashedName: step-57
---

# --description--

حاليًا، تُشفّر كل حرف دائمًا بنفس الحرف، اعتمادًا على الإزاحة المحددة. ماذا لو كانت الإزاحة مختلفة لكل حرف؟ سيكون فك التشفير أصعب بكثير. تُسمى هذه الخوارزمية تشفير فيجينير، حيث تُحدد الإزاحة لكل حرف بواسطة نص آخر يُسمى *المفتاح*.

ابدأ بتحويل تشفير قيصر إلى تشفير فيجينير بإزالة استدعائي الدالة.

# --hints--

يجب إزالة استدعائي الدالة `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
