---
id: 6553a7d8d05cbb1ae335a665
title: الخطوة 45
challengeType: 20
dashedName: step-45
---

# --description--

عندما تصل الحلقة إلى الحرف `Z`، يتجاوز مجموع `index + shift` آخر فهرس في سلسلة النص `alphabet`. لذلك، تحاول `alphabet[new_index]` استخدام فهرس غير صالح، مما يسبب رمي `IndexError`.

يمكنك ملاحظة أن الإخراج في الطرفية يتوقف عند الفراغ مباشرة قبل `Z`، وكان آخر `print` قبل رمي الخطأ.

في هذه الحالة، يمكن استخدام عامل باقي القسمة (`%`) لإرجاع الباقي من قسمة رقمين. على سبيل المثال: `5 % 2` يساوي `1`، لأن 5 مقسومة على 2 حاصل قسمتها 2 وباقي قسمتها 1.

أحط `index + shift` بـأقواس عادية، وطبق باقي القسمة على التعبير بـ `26`، وهو طول `alphabet`.

# --hints--

يجب أن يكون لديك `new_index = (index + shift) % 26` في جملة `else` الخاصة بك.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
