---
id: 6554a88d5af937226f4a9121
title: الخطوة 65
challengeType: 20
dashedName: step-65
---

# --description--

طريقة `.index()` مطابقة لطريقة `.find()` لكنها ترمي استثناء `ValueError` إذا لم تستطع إيجاد السلسلة النصية الفرعية.

الاستثناء `ValueError` هو استثناء مدمج يُرمي عندما تُمرر معلمة ذات نوع صحيح لكن بقيمة غير مناسبة إلى دالة.

بعد زيادة `key_index`، أعلن متغيرًا باسم `offset`. ابحث عن الفهرس الذي تحمله `key_char` في `alphabet` وعيّنه إلى `offset`. استخدم `.index()` لإيجاد الفهرس.

# --hints--

يجب إعلان متغير يسمى `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

يجب أن يخزن متغير `offset` قيمة `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
