---
id: 65549561463f0016876e852c
title: الخطوة 60
challengeType: 20
dashedName: step-60
---

# --description--

بما أن مفتاحك أقصر من النص الذي ستحتاج إلى تشفيره، يجب تكراره لتوليد النص المشفر بالكامل.
في بداية جسم الدالة، أعلن عن متغير `key_index` وعيّنه إلى صفر.

# --hints--

يجب إعلان متغير باسم `key_index` في بداية جسم الدالة.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

يجب تعيين القيمة `0` إلى متغير `key_index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تنسيق غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
