---
id: 6554ad2463b8892748f8efdd
title: الخطوة 67
challengeType: 20
dashedName: step-67
---

# --description--

في الوقت الحالي، تقوم الدالة بطباعة بعض السلاسل النصية، لكن لا يمكن للأجزاء الأخرى من الشفرة استخدام هذه القيم لأداء أي عمليات.

لهذا الغرض، تحتاج إلى استخدام جملة `return`:

```python
def foo():
    return 'spam'
```

عليك كتابة `return` متبوعة بمسافة والقيمة التي يجب أن تُرجعها الدالة. بمجرد العثور على جملة `return`، تُرجع تلك القيمة ويتوقف تنفيذ الدالة، ثم ينتقل التنفيذ إلى السطر التالي من الشفرة بعد استدعاء الدالة. في المثال أعلاه، تُرجع الدالة `foo` السلسلة النصية `'spam'`.

احذف استدعائي `print()` من دالتك وأرجع `encrypted_text`.

# --hints--

يجب حذف استدعائي `print()` من دالتك.

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

يجب أن تُرجع دالتك `encrypted_text`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
