---
id: 6554a334a40edb1fb4eff827
title: الخطوة 62
challengeType: 20
dashedName: step-62
---

# --description--

بعد ذلك، داخل كتلة `else`، أعلن متغيرًا يسمى `key_char` وقم بتعيين القيمة من `key` عند الفهرس `key_index` مع باقي قسمة (`%`) طول `key`.

# --hints--

يجب إعلان متغير يسمى `key_char` في بداية جسم بند `else` الخاص بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/^\s*key_char\s*=/));
```

يجب تعيين `key[key_index % len(key)]` إلى متغير `key_char` الخاص بك.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");    
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]/));
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--        
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
--fcc-editable-region--        
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
