---
id: 6554e0adc7bb193cbfdb36d5
title: الخطوة 73
challengeType: 20
dashedName: step-73
---

# --description--

تحقق مما إذا كانت الدالة تستطيع فك تشفير السلسلة النصية إلى النص العادي.

أعلن عن متغير آخر اسمه `decryption` وقم بتعيينه إلى `vigenere(encryption, custom_key, -1)`.

# --hints--

يجب أن تستدعي `vigenere` مع تمرير `encryption` و`custom_key` و`-1` كمعلمات.

```js
({ test: () => assert.match(code, /vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/) })
```

يجب أن تعلن عن متغير `decryption`.

```js
({ test: () => assert(__userGlobals.has("decryption")) })
```

يجب أن يحتوي متغير `decryption` على القيمة `vigenere(encryption, custom_key, -1)`.

```js
({ test: () => assert.match(code, /^decryption\s*=\s*vigenere\s*\(\s*encryption\s*,\s*custom_key\s*,\s*-\s*1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key, direction):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
--fcc-editable-region--    
encryption = vigenere(text, custom_key, 1)
print(encryption)

--fcc-editable-region--
```
