---
id: 6554930320d70414e7b6acc6
title: الخطوة 59
challengeType: 20
dashedName: step-59
---

# --description--

احذف متغير `shift`. ثم أعلن متغيرًا جديدًا باسم `custom_key` وقم بتعيين قيمة السلسلة النصية `'python'` لهذا المتغير.

# --hints--

يجب حذف المتغير `shift` وقيمته.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

يجب إعلان متغير باسم `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

يجب تعيين السلسلة النصية `'python'` إلى متغير `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
