---
id: 655491bd5b98b813fa5bedca
title: الخطوة 58
challengeType: 20
dashedName: step-58
---

# --description--

عدّل إعلان الدالة الخاص بك: غيّر اسم الدالة إلى `vigenere` والمعلمة الثانية إلى `key`.

# --hints--

يجب أن تعدل اسم دالتك إلى `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

يجب أن تأخذ دالة `vigenere` المعلمتين `message` و `key`.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
