---
id: 65549f90cf78131c96ebcf28
title: الخطوة 61
challengeType: 20
dashedName: step-61
---

# --description--

عند كتابة الشفرة، تُعد قابلية القراءة أمرًا أساسيًا. التعليقات تعمل كملاحظات فعالة تشرح المنطق وراء الشفرة. تكون ذات قيمة عند العودة إلى مشروع بعد فترة، وتساعد زملاءك أيضًا على فهم الشفرة.

في Python، يمكنك كتابة تعليق باستخدام `#`. أي شيء يأتي بعد `#` لن يُنفذ.

قبل جملة `if`، أضف تعليقًا يقول `Append space to the message`.

# --hints--

يجب أن تضيف تعليقًا يقول `Append space to the message`. لا تنسَ رمز `#` في البداية.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
