---
id: 6555db4a7b788e15795674e5
title: الخطوة 89
challengeType: 20
dashedName: step-89
---

# --description--

يمكن دمج سلسلتين نصيتين أو أكثر باستخدام العامل `+`. على سبيل المثال: `'Hello' + ' there!'` ينتج `'Hello there!'`.

استدعِ الدالة `print()` واستخدم العامل `+` لدمج المتغير `text` مع السلسلة النصية `'Encrypted text: '`. انتبه إلى الفراغات.

# --hints--

يجب استدعاء `print()` وتمرير `'Encrypted text: ' + text` كمعلمات.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^print\s*\(\s*("|')Encrypted\stext:\s\1\s*\+\s*text\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--    

--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
