---
id: 6555e1bfcd374e18c6be8e58
title: الخطوة 91
challengeType: 20
dashedName: step-91
---

# --description--

في بايثون، توجد طريقة سهلة لتنسيق السلاسل النصية. تتيح لك *f-strings* استيفاء القيم داخل سلاسل النص.

الاستيفاء يعني كتابة عناصر نائب تُستبدَل بالقيم المحددة عند تشغيل البرنامج. على سبيل المثال، يمكنك الحصول على نفس نتيجة `'Encrypted text: ' + text` باستخدام `f'Encrypted text: {text}'`. يجب وضع `f` قبل علامات الاقتباس لإنشاء f-string وكتابة المتغيرات أو التعبيرات التي تريد استيفاءها بين الأقواس المتعرجة.

عدّل أول استدعاء لـ `print()` لاستخدام f-string.

# --hints--

يجب تعديل استدعاء `print('Encrypted text: ' + text)` ليطبع نفس النتيجة باستخدام f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
