---
id: 6555d6e12c95701172b55709
title: الخطوة 77
challengeType: 20
dashedName: step-77
---

# --description--

يمكن استدعاء الدوال بمعلمات افتراضية. تشير المعلمة الافتراضية إلى القيمة التي يجب أن تأخذها الدالة إذا لم تُمرر المعلمة. على سبيل المثال، الدالة أدناه تقبل ثلاث معلمات لكن يمكنك استدعاؤها بتمرير معلمتين فقط. ستفترض المعلمة الثالثة القيمة المحددة افتراضيًا:

```py
def foo(a, b, c=value):
    <code>
```

عدّل دالة `vigenere` بحيث تكون للمعلمة `direction` قيمة افتراضية تساوي `1`.

# --hints--

يجب أن تكون للمعلمة `direction` في دالة `vigenere` قيمة افتراضية تساوي `1`.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
