---
id: 6555231eac4d9f19bd3d44b5
title: الخطوة 83
challengeType: 20
dashedName: step-83
---

# --description--

يمكن استخدام الكلمة المفتاحية `pass` كعنصر نائب لشفرات مستقبلية. لا تؤثر على شفرتك لكنها تحميك من الأخطاء التي قد تحدث في حالة وجود شفرة غير مكتملة:

```py
def foo():
    pass
```

استدعاء `vigenere` مع `1` للتشفير و `-1` لفك التشفير مقبول لكنه قد يكون غامضًا بعض الشيء. أنشئ دالة جديدة باسم `encrypt` تأخذ معلمتي `message` و `key`، واستخدم `pass` لملء جسم الدالة.


# --hints--

يجب إعلان دالة جديدة باسم `encrypt`. تذكر استخدام الكلمة المفتاحية `pass` داخل الجسم.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

يجب أن تأخذ دالة `encrypt` معلمتي `message` و `key`. ترتيب المعلمات مهم.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
