---
id: 655a2a7210094920069b117c
title: الخطوة 40
challengeType: 20
dashedName: step-40
---

# --description--

تسمح لك معاملات المقارنة بمقارنة كائنين بناءً على قيمتهما. يمكنك استخدام معامل مقارنة بوضعه بين الكائنين اللذين تريد مقارنتهما.
تعيد قيمة *منطقية* — وهي `True` أو `False` — اعتمادًا على صحة التعبير.

يحتوي Python على معاملات المقارنة التالية:

<table>
  <thead>
    <tr>
<th>المعامل</th>
      <th>المعنى</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
<td>مساوٍ</td>
    </tr>
    <tr>
      <td>!=</td>
<td>غير مساوي</td>
    </tr>
    <tr>
<td>&gt;</td>
      <td>أكبر من</td>
    </tr>
    <tr>
<td>&lt;</td>
      <td>أصغر من</td>
    </tr>
    <tr>
<td>&gt;=</td>
      <td>أكبر من أو يساوي</td>
    </tr>
    <tr>
<td>&lt;=</td>
      <td>أصغر من أو يساوي</td>
    </tr>
  </tbody>
</table>

في بداية جسم الحلقة، اطبع نتيجة مقارنة `char` بمسافة (`' '`). استخدم معامل المساواة `==` لذلك.

# --hints--

يجب أن تقارن `char` بمسافة باستخدام معامل المساواة داخل حلقة `for` الخاصة بك.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

يجب أن تطبع نتيجة مقارنة `char` بمسافة داخل حلقة `for` الخاصة بك.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

يجب أن تطبع نتيجة مقارنة `char` بمسافة في بداية الحلقة.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
