---
id: 6568b7a40c7ba0ccbc4b4425
title: الخطوة 32
challengeType: 20
dashedName: step-32
---

# --description--

كما ترى، تُطبع كل حرف من السلسلة النصية المُعيَّنة داخل الحلقة.

ارجع إلى السلسلة النصية الأصلية بحذف تعيين `text` الجديد.

# --hints--

يجب ألا يكون لديك `text = 'Albatross'` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*=\s*("|')Albatross\1/.test(commentless_code))
```

يجب أن تظل لديك متغير `text` بقيمة `'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code, /text\s*=\s*("|')Hello World\1/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text = 'Albatross'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
