---
id: 65687d2f8c7ee27b0446cef3
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

الكود الذي يُنفذ في كل تكرار — الموضوع بعد `:` — يشكل جسم الحلقة. يجب أن يكون هذا الكود مُزاحًا بمسافة بادئة. في بايثون، يُنصح باستخدام 4 مسافات لكل مستوى من مستويات المسافة البادئة. هذا المستوى المزاح هو كتلة كود.

```py
for i in text:
    print(i)
```

تعتمد بايثون على المسافة البادئة لتحديد كتل الكود. النقطتان الرأسيتان في نهاية السطر تشير إلى أن كتلة كود جديدة مزاحة ستتبع.

لذا، عندما لا توجد كتلة مزاحة بعد النقطتين الرأسيتين الأخيرتين، يتوقف تنفيذ الكود ويُرمى خطأ `IndentationError`. هذا الكود لن يعرض الناتج وبدلاً من ذلك سيرجع خطأ `IndentationError`:

```py
for i in text:
print(i)
```

أضف جسمًا لحلقة `for` الخاصة بك بإضافة استدعاء إلى `print(i)`. تذكر أن تُدخل جسم الحلقة.

# --hints--

يجب أن تضيف `print(i)` إلى جسم حلقة `for` الخاصة بك. انتبه إلى التداخل.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

يحتوي كودك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
