---
id: 656886ec4e4d9b96ded3c565
title: الخطوة 25
challengeType: 20
dashedName: step-25
---

# --description--

يمكن أن يحمل متغير التكرار أي اسم صالح، لكن من الملائم أن تعطيه اسمًا ذا معنى.

أعد تسمية متغير `i` إلى `char`.

# --hints--

يجب أن تستبدل `i` بـ `char`.

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تنسيق غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
