---
id: 65687d2f8c7ee27b0446cef3
title: الخطوة 24
challengeType: 20
dashedName: step-24
---

# --description--

الشفرة التي تُنفذ في كل تكرار — الموضوعة بعد `:` — تشكل جسم الحلقة. يجب أن تكون هذه الشفرة متداخلة. في Python، يُنصح باستخدام 4 مسافات لكل مستوى من التداخل. هذا المستوى المتداخل هو كتلة شفرة.

```py
for i in text:
    print(i)
```

تعتمد Python على التداخل للدلالة على كتل الشفرة. النقطتان الرأسيتان في نهاية السطر تشير إلى أن كتلة شفرة متداخلة جديدة ستتبع.

لذا، عندما لا توجد كتلة متداخلة بعد النقطتين الرأسيتين الأخيرتين، يتوقف تنفيذ الشفرة ويُرمى `IndentationError`. هذه الشفرة لن تعرض الناتج وبدلاً من ذلك سترمي `IndentationError`:

```py
for i in text:
print(i)
```

أضف جسمًا لحلقة `for` الخاصة بك بإضافة استدعاء إلى `print(i)`. تذكر أن تُدخل جسم الحلقة.

# --hints--

يجب أن تضيف `print(i)` إلى جسم حلقة `for` الخاصة بك. انتبه إلى التداخل.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تداخل غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
