---
id: 65688b5a1655a7a6caede847
title: الخطوة 31
challengeType: 20
dashedName: step-31
---

# --description--

عندما تحاول تغيير الأحرف الفردية في سلسلة نصية كما فعلت في الخطوة السابقة، تحصل على `TypeError`، الذي يحدث عندما يُستخدم كائن من نوع غير مناسب في شفرتك.

كما ترى من رسالة الخطأ، السلاسل النصية لا تدعم تعيين العناصر، لأنها غير قابلة للتغيير. مع ذلك، يمكن إعادة تعيين متغير إلى سلسلة نصية أخرى:

```py
message = 'Hello World'
message = 'Hello there!'
```

احذف السطر `text[0]` وأعد تعيين `text` إلى السلسلة النصية `'Albatross'`.

# --hints--

يجب ألا يحتوي شفرتك على `text[0] = 'C'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

يجب إعادة تعيين `text` إلى السلسلة النصية `'Albatross'`. لا تعدل `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
