---
id: 65688a50e6c998a21d8e41d3
title: الخطوة 30
challengeType: 20
dashedName: step-30
---

# --description--

السلاسل النصية غير قابلة للتغيير، مما يعني أنه لا يمكن تعديلها بعد إنشائها. على سبيل المثال، قد تظن أن الشفرة التالية تغير قيمة `my_string` إلى السلسلة النصية `'train'`، لكن هذا غير صحيح:

```py
my_string = 'brain'
my_string[0] = 't'
```

تأكد من استخدام ترميز الأقواس للوصول إلى الحرف الأول في `text` وحاول تغييره إلى حرف تختاره. سترى أن الناتج يختفي ويظهر خطأ.

# --hints--

يجب أن يبقى لديك `text = 'Hello World'` في شفرتك.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

يجب أن تصل إلى الحرف الأول في `text` باستخدام `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

يجب أن تستخدم عامل `=` لتعيين حرف تختاره إلى `text[0]`. لا تنسَ إحاطة الحرف بعلامات اقتباس مفردة أو علامات اقتباس مزدوجة.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
