---
id: 65951b637ddec63611c482c0
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

يمكنك استخدام الدالة المدمجة `print()` لطباعة ناتج الشفرة على الطرفية.

الدوال هي كتل شفرة قابلة لإعادة الاستخدام يمكنك استدعاؤها لتشغيل شفرتها عند الحاجة. لاستدعاء دالة، يكفي كتابة زوج من الأقواس العادية بجانب اسمها. ستتعلم المزيد عن الدوال قريبًا جدًا.

في الوقت الحالي، انتقل إلى سطر جديد وأضف استدعاء فارغ للدالة `print()`. لا يجب أن ترى أي ناتج حتى الآن.

# --hints--

يجب أن تحتوي شفراتك على `print()`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

تحتوي شفراتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
