---
id: 659526f3e5424b3eac4b6aa7
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

بعد ذلك، اطبع المتغير `index` في وحدة التحكم.

# --hints--

يجب أن تطبع المتغير `index`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تنسيق غير صحيح للمسافات البادئة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
