---
id: 65952b70924edb449faf1950
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

أعلن عن متغير جديد اسمه `shifted`. استخدم ترميز الأقواس للوصول إلى قيمة `alphabet` عند الفهرس `index` وقم بتعيينها إلى المتغير الجديد.

# --hints--

يجب إعلان متغير اسمه `shifted`. انتبه لوضع اسم المتغير في بداية السطر.

```js
assert.match(code, /^shifted\s*=/m)
```

يجب تعيين `alphabet[index]` إلى متغير `shifted` الخاص بك.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

تحتوي شفرتك على صيغة نحوية غير صحيحة و/أو تراجع غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
