---
id: 659530ca823d6c47f61daa8a
title: الخطوة 20
challengeType: 20
dashedName: step-20
---

# --description--

اطبع المتغير `shifted` الخاص بك.

# --hints--

يجب أن تطبع المتغير `shifted`. انتبه لوضع استدعاء الدالة في بداية السطر.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

يحتوي كودك على صيغة نحوية غير صحيحة و/أو مسافة بادئة غير صحيحة.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
