---
id: 659db3e3670d3e712be82593
title: الخطوة 22
challengeType: 20
dashedName: step-22
---

# --description--

تكرار عملية تحديد موقع الحرف داخل الأبجدية وتحديد الحرف المنقول لكل رمز في `text` قد يكون مرهقًا. لحسن الحظ، يمكنك تبسيط ذلك باستخدام حلقة.

في الوقت الحالي، احذف كل أسطر الشفرة أسفل إعلان المتغير `alphabet`.

# --hints--

يجب أن يبقى لديك `text = 'Hello World'` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

يجب أن يبقى لديك `shift = 3` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

يجب أن يبقى لديك `alphabet = 'abcdefghijklmnopqrstuvwxyz'` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

يجب حذف المتغير `index` وقيمته.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

يجب ألا يكون لديك `print(index)` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

يجب حذف المتغير `shifted` وقيمته.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

يجب ألا يكون لديك `print(shifted)` في شفرتك.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

شفرتك تحتوي على صيغة نحوية غير صحيحة و/أو تنسيق مسافات بادئة غير صحيح.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
