---
id: 6624e0ec0bf47017eed84b1f
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

النوع الأول من الشيفرات الذي ستبنيه يُسمى شيفرة *قيصر*. بالتحديد، ستأخذ كل حرف في رسالتك، تجد موقعه في الأبجدية، تأخذ الحرف الموجود بعد 3 مواقع في الأبجدية، وتستبدل الحرف الأصلي بالحرف الجديد.

لتنفيذ ذلك، ستستخدم طريقة `.find()` التي نوقشت في الخطوة السابقة. عدل استدعاء `.find()` الحالي بتمرير `text[0]` كمعلمة بدلاً من `'z'`.

# --hints--

يجب تعديل استدعاء `alphabet.find('z')` الحالي بتمرير `text[0]` إلى الطريقة.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
