---
id: 66605a64a9822f01ec5ae8ce
title: الخطوة 50
challengeType: 20
dashedName: step-50
---

# --description--

في Python، يحدد *نطاق* المتغير المكان الذي يمكن الوصول إليه فيه:

- المتغيرات المعرفة خارج الدالة لها نطاق *عالمي* ويمكن الوصول إليها من أي جزء في شفرتك.

- المتغيرات المعرفة داخل الدالة تكون *محلية* لتلك الدالة ولا يمكن الوصول إليها من خارجها.

لتشاهد هذا عمليًا، حاول طباعة المتغير `alphabet` في نهاية شفرتك. سيؤدي ذلك إلى رفع استثناء `NameError`.

يجب أن ترى رسالة خطأ تشير إلى أن `alphabet` غير معرف. هذا لأن `alphabet` معرف داخل الدالة `caesar` ولا يمكن الوصول إليه من خارجها.

# --hints--

يجب أن تحاول طباعة المتغير `alphabet` خارج دالة caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
