---
id: 66cc172af7159a7c67804544
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Sasa ni wakati wa kuongeza utendakazi wa kuzidisha kwenye kalkuleta yako.

Tangaza kitendakazi kinachoitwa `calculateProduct` kinachopokea vigezo viwili, `num1` na `num2`.

Ndani ya kitendakazi hicho, rudisha bidhaa ya vigezo viwili.

Kisha, ongeza `console.log` inayofanya wito wa kitendakazi cha `calculateProduct` kwa hoja `13` na `5`.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

Kitendakazi chako cha `calculateProduct` kinapaswa kuwa na vigezo viwili vinavyoitwa `num1` na `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

Kitendakazi chako cha `calculateProduct` kinapaswa kurudisha bidhaa ya vigezo viwili.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateProduct` kwa hoja `13` na `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
