---
id: 66cc0f1ae40802781b2ea972
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na vigezo vya kitendakazi.

Kigezo cha kitendakazi ni kigezo kinachotangazwa katika tamko la kitendakazi na hufanya kazi kama kiwakilishi.

Hapa kuna mfano wa kitendakazi chenye kigezo:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Tangaza kitendakazi kinachoitwa `calculateSum` kinachopokea vigezo viwili, `num1` na `num2`.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `calculateSum`.

```js
assert.isFunction(calculateSum);
```

Kitendakazi chako cha `calculateSum` kinapaswa kupokea vigezo viwili, `num1` na `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
