---
id: 66cc0ba5881acb7692cfc4de
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa ni wakati wa kuongeza kitendakazi kingine.

Tangaza kitendakazi kinachoitwa `addThreeAndFour` kinachorejesha jumla ya `3` na `4`.

Kisha fanya wito wa kitendakazi cha `addThreeAndFour` ndani ya `console.log` kuona matokeo.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Kitendakazi chako cha `addThreeAndFour` kinapaswa kurudisha jumla ya `3` na `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

Kitendakazi chako cha `addThreeAndFour` kinapaswa kurudisha jumla kwa kutumia mtaalamu wa `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

Unapaswa kufanya wito wa kitendakazi cha `addThreeAndFour` ndani ya `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
