---
id: 66cc095dfe1682753d2ab030
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Moja ya dhana ulizojifunza katika mafundisho yaliyopita ilikuwa jinsi ya kurudisha thamani kutoka kwenye kitendakazi.

Hapa kuna ukumbusho wa jinsi ya kurudisha thamani kutoka kwenye kitendakazi:

```js
function myFunction() {
  return "Hello World";
}
```

Ndani ya kitendakazi chako cha `addTwoAndSeven`, rudisha jumla ya `2` na `7`.

# --hints--

Kitendakazi chako cha `addTwoAndSeven` kinapaswa kurudisha jumla ya `2` na `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Kitendakazi chako cha `addTwoAndSeven` kinapaswa kurudisha jumla kwa kutumia kiendeshaji cha `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
