---
id: 66cc16abc30b367ba1d1f9ef
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako cha `calculateDifference`.

Anza kwa kuongeza `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `22` na `5`.

Kisha, ongeza `console.log` ya pili inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `12` na `1`.

Mwishowe, ongeza `console.log` ya tatu inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `17` na `9`.

# --hints--

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `22` na `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `12` na `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `17` na `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
