---
id: 66cc281d687975858049fd8d
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Katika hatua ya mwisho ya warsha, utahitaji kujaribu kitendakazi chako cha `calculateSquareRoot`.

Ongeza `console.log` ya kwanza inayofanya wito wa kitendakazi cha `calculateSquareRoot` kwa hoja ya `25`.

Kisha, ongeza `console.log` ya pili inayofanya wito wa kitendakazi cha `calculateSquareRoot` kwa hoja ya `100`.

Na kwa mabadiliko hayo, programu yako ya kalkuleta itakuwa imekamilika!

# --hints--

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateSquareRoot` kwa hoja ya `25`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquareRoot\s*\(\s*25\s*\)\s*\)/);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateSquareRoot` kwa hoja ya `100`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquareRoot\s*\(\s*100\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

function calculateSquareRoot(num) {
  return Math.sqrt(num);
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

function calculateSquareRoot(num) {
  return Math.sqrt(num);
}

console.log(calculateSquareRoot(25));
console.log(calculateSquareRoot(100));
```
