---
id: 66cc12fa504b0479dac479a0
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Faida ya vitendakazi ni kwamba vinaweza kuitwa na hoja tofauti, ikikuruhusu kutumia msimbo uleule tena na thamani mbalimbali.

Hapa kuna mifano ya miito ya vitendakazi na hoja tofauti za mfuatano wa herufi:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Ongeza `console.log` inayoitwa kitendakazi cha `calculateSum` na hoja `2` na `5`.

# --hints--

Unapaswa kuwa na `console.log` inayoitwa kitendakazi cha `calculateSum` na hoja `2` na `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
