---
id: 66cc2739f687e484d50bb6f1
title: Hatua 18
challengeType: 1
dashedName: step-18
---

# --description--

Sehemu ya mwisho ya kitendakazi cha kalkuleta yako itakuwa kitendakazi cha mzizi wa mraba.

Mraba ni nambari inayojumlishwa na yenyewe. Hivyo, mzizi wa mraba utakuwa nambari ambayo inapojumlishwa na yenyewe hutoa nambari asili.

Kwa mfano, mzizi wa mraba wa `9` ni `3` kwa sababu `3 * 3 = 9`.

Ili kupata mzizi wa mraba wa nambari katika JavaScript, unaweza kutumia njia ya <dfn>Math.sqrt()</dfn>.

Hapa kuna mfano:

```js
// result: 3
Math.sqrt(9);
```

Tangaza kitendakazi kinachoitwa `calculateSquareRoot` kinachopokea kigezo cha `num` na kurudisha mzizi wa mraba wa `num`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Kitendakazi chako cha `calculateSquareRoot` kinapaswa kupokea kigezo cha `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi chako cha `calculateSquareRoot` kinapaswa kurudisha mzizi wa mraba wa `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
