---
id: 66cc21d23238dc8240a8a182
title: Hatua 16
challengeType: 1
dashedName: step-16
---

# --description--

Itakuwa vizuri ikiwa kalkuleta yako itahesabu mraba wa nambari. Mraba wa nambari ni nambari inayozidishwa na yenyewe.

Ili kuhesabu mraba wa nambari katika JavaScript, unaweza kutumia njia ya <dfn>Math.pow()</dfn>. Njia ya `Math.pow()` inachukua hoja mbili: nambari ya msingi na nguvu.

Pia unaweza kutumia kigezo cha nguvu (`**`) kuhesabu mraba wa nambari.

Hapa kuna mfano:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Tangaza kitendakazi kinachoitwa `calculateSquare` kinachochukua kigezo cha `num` na kurudisha mraba wa `num`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

Kitendakazi chako cha `calculateSquare` kinapaswa kuchukua kigezo cha `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi chako cha `calculateSquare` kinapaswa kurudisha mraba wa `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
