---
id: 66cc06012e46aa6bc9b8c001
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utakagua jinsi ya kufanya kazi na vitendakazi kwa kujenga programu ya kalkuleta.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kutangaza vitendakazi kama hivi:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Anza kwa kutengeneza kitendakazi kinachoitwa `addTwoAndSeven`.

Unaweza kuchagua kutumia sintaksia ya kitendakazi cha kawaida au sintaksia ya kitendakazi cha mshale.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
