---
id: 66cc0de2920bec775f610424
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Ingawa vitendakazi vinafanya kazi kama inavyotarajiwa, kuna rudia nyingi katika msimbo wako.

Kwa kuwa unajenga kalkuleta, hutaki kuunda kitendakazi kwa kila mchanganyiko unaowezekana wa nambari unazotaka kuongeza pamoja.

Itakuwa bora kuunda kitendakazi kimoja kinachoweza kutumika tena ambacho kinaweza kuongeza nambari mbili yoyote pamoja.

Ondoa msimbo wote uliouandika hadi sasa.

Katika hatua inayofuata, utakagua jinsi ya kufanya kazi na vigezo na hoja katika vitendakazi.

# --hints--

Hautaki kuwa na kitendakazi kinachoitwa `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Hautaki kuwa na wito wa kitendakazi kwa `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Hautaki kuwa na kitendakazi kinachoitwa `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Hautaki kuwa na wito wa kitendakazi kwa `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
