---
id: 66cc0a9e06e00b75d6782be9
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kuita (kuitisha) kitendakazi. Kuita kitendakazi kunamaanisha kutekeleza msimbo ulioko ndani ya kitendakazi.

Hapa kuna ukumbusho wa jinsi ya kuita kitendakazi:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Ongeza tamko la `console.log`, na ndani yake, ita kitendakazi cha `addTwoAndSeven`.

Sasa unapaswa kuona jumla ya `2` na `7` ikionyeshwa kwenye konsoli.

# --hints--

Unapaswa kuwa na `console.log` katika msimbo wako.

```js
assert.match(code, /console\.log(.*)/);
```

Unapaswa kuita kitendakazi cha `addTwoAndSeven` ndani ya `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
