---
id: 66cc16abc30b367ba1d1f9ef
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Sasa ni wakati wa kujaribu kitendakazi chako cha `calculateDifference`.

Anza kwa kuongeza `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `22` na `5`.

Kisha, ongeza `console.log` ya pili inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `12` na `1`.

Mwishowe, ongeza `console.log` ya tatu inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `17` na `9`.

# --hints--

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `22` na `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `12` na `1`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi cha `calculateDifference` kwa hoja `17` na `9`.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
