---
id: 66cc1a3a39aef47d6473cb2f
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Hatua inayofuata ni kuongeza utendaji wa kugawanya kwenye kalkuleta yako.

Tangaza kitendakazi kinachoitwa `calculateQuotient` kinachopokea vigezo viwili, `num1` na `num2`.

Ndani ya kitendakazi hicho, rudisha gawio la vigezo viwili.

Kisha, ongeza `console.log` inayofanya wito wa kitendakazi `calculateQuotient` kwa hoja `7` na `11`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

Kitendakazi chako cha `calculateQuotient` kinapaswa kuwa na vigezo viwili vinavyoitwa `num1` na `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi chako cha `calculateQuotient` kinapaswa kurudisha gawio la vigezo viwili.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Unapaswa kuwa na `console.log` inayofanya wito wa kitendakazi `calculateQuotient` kwa hoja `7` na `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
