---
id: 66cc0ba5881acb7692cfc4de
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa, ni wakati wa kuongeza kitendakazi kingine.

Tangaza kitendakazi kinachoitwa `addThreeAndFour` kinachorejesha jumla ya `3` na `4`.

Kisha fanya wito wa kitendakazi cha `addThreeAndFour` ndani ya `console.log` kuona matokeo.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

Kitendakazi chako cha `addThreeAndFour` kinapaswa kurudisha jumla ya `3` na `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

Kitendakazi chako cha `addThreeAndFour` kinapaswa kurudisha jumla kwa kutumia mtaalamu wa `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Unapaswa kufanya wito wa kitendakazi cha `addThreeAndFour` ndani ya `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
