---
id: 66cc10958d72e57907f898b2
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ndani ya kitendakazi chako cha `calculateSum`, utahitaji kurudisha jumla ya vigezo viwili.

# --hints--

Kitendakazi chako cha `calculateSum` kinapaswa kurudisha jumla ya vigezo viwili.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
