---
id: 66cc14ca14e65e7a9ebecb80
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Sasa baada ya kalkuleta yako kuweza kuongeza nambari mbili pamoja, ni wakati wa kutangaza kitendakazi kitakachopunguza nambari mbili.

Tangaza kitendakazi kinachoitwa `calculateDifference` kinachopokea vigezo viwili, `num1` na `num2`.

Ndani ya kitendakazi, rudisha tofauti kati ya vigezo viwili.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

Kitendakazi chako cha `calculateDifference` kinapaswa kuwa na vigezo viwili vinavyoitwa `num1` na `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi chako cha `calculateDifference` kinapaswa kurudisha tofauti kati ya vigezo viwili.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
