---
id: 66cdf76685e4cb5a8726e27b
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Ili kujaribu zaidi kitendakazi chako cha `calculateSum`, utahitaji kukiwita kwa hoja tofauti.

Ongeza wito mwingine wa `console.log` unaoitisha kitendakazi cha `calculateSum` kwa hoja `10` na `10`.

Chini ya wito huo wa `console.log`, ongeza wito mwingine wa `console.log` unaoitisha kitendakazi cha `calculateSum` kwa hoja `5` na `5`.

# --hints--

Unapaswa kuwa na wito wa `console.log` unaoitisha kitendakazi cha `calculateSum` kwa hoja `10` na `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

Unapaswa kuwa na wito wa `console.log` unaoitisha kitendakazi cha `calculateSum` kwa hoja `5` na `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
