---
id: 63507e4562cdde3a28e8de1b
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

تمامًا مثل طريقة `.map()`، تأخذ طريقة `.filter()` دالة استدعاء راجع. تأخذ دالة الاستدعاء الراجع العنصر الحالي كمعلمتها الأولى.

```js
array.filter(el => {

})
```

يجب أن تُرجع دالة الاستدعاء الراجع قيمة منطقية، تشير إلى ما إذا كان يجب تضمين العنصر في المصفوفة الجديدة. في هذه الحالة، تريد إرجاع `true` إذا لم يكن العنصر `NaN` (ليس رقمًا).

مع ذلك، لا يمكنك التحقق من المساواة هنا، لأن `NaN` لا يساوي نفسه. بدلًا من ذلك، يمكنك استخدام طريقة `isNaN()`، التي تُرجع `true` إذا كانت المعلمة `NaN`.

أضف دالة استدعاء راجع إلى طريقة `.filter()` مع `el` كمعلمة، واستخدم الإرجاع الضمني لإرجاع `true` إذا لم يكن العنصر `NaN`.

# --hints--

يجب أن تحتوي طريقة `.filter()` على دالة استدعاء راجع تقبل `el` كمعلمة.

```js
assert.match(calculate.toString(), /numbers\.filter\(\s*(\(\s*el\s*\)|el)\s*=>|numbers\.filter\(\s*function\s*\(\s*el\s*\)\s*\{/)
```

يجب أن تستخدم دالة الاستدعاء الراجع `!` و`isNaN()` للتحقق مما إذا كان `el` ليس `NaN`.

```js
assert.match(calculate.toString(), /!\s*(Number\.)?isNaN\(\s*el\s*\)/);
```

يجب أن تُرجع دالة الاستدعاء الراجع العناصر التي ليست `NaN`.

```js
assert.match(calculate.toString(), /(numbers\.filter\(\s*(\(\s*el\s*\)|el)\s*=>|numbers\.filter\(\s*function\s*\(\s*el\s*\)\s*\{)\s*(return\s*)?!(Number\.)?isNaN\(\s*el\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el));
  const filtered = numbers.filter();
}
--fcc-editable-region--
```
