---
id: 63507d810f1a2e38f1908fd8
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

يمكن للمستخدم إدخال أي نص يريده في مربع المُدخَل. تريد التأكد من أنك تتعامل فقط مع الأرقام. دالة البناء `Number()` سترجع `NaN` (التي تعني "ليست رقمًا") إذا لم يكن بالإمكان تحويل القيمة الممررة إليها إلى رقم.

تحتاج إلى تصفية هذه القيم – ولحسن الحظ، المصفوفات تحتوي على طريقة مخصصة لذلك. طريقة `.filter()` تسمح لك بتصفية العناصر من المصفوفة، مع إنشاء مصفوفة جديدة في العملية.

أعلن عن متغير `filtered` وقم بتعيين `numbers.filter()` له.

# --hints--

يجب أن تحتوي دالة `calculate` على متغير `filtered`.

```js
assert.match(calculate.toString(), /filtered/);
```

يجب أن تستخدم دالة `calculate` طريقة `.filter()` على مصفوفة `numbers`.

```js
assert.match(calculate.toString(), /numbers\.filter\(\s*\)/)
```

يجب أن تعيّن نتيجة `numbers.filter()` إلى المتغير `filtered`.

```js
assert.match(calculate.toString(), /filtered\s*=\s*numbers\.filter\(\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>

    <form onsubmit="calculate();">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
--fcc-editable-region--
const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el));
}
--fcc-editable-region--
```
