---
id: 661899113600ef307d8a7487
title: الخطوة 23
challengeType: 0
dashedName: step-23
---

# --description--

للحصول على الوسيط لمصفوفة تحتوي على عدد فردي من العناصر، عليك إيجاد الرقم الأوسط وإرجاعه.

إليك كيفية إيجاد الرقم الأوسط لمصفوفة تحتوي على عدد فردي من العناصر:

```js
arr[Math.floor(arr.length / 2)];
```

إليك مثالًا أطول لإيجاد الرقم الأوسط لمصفوفة تحتوي على `5` عناصر:

```js
const numbers = [1, 2, 3, 4, 5];
const middleNumber = numbers[Math.floor(numbers.length / 2)];
console.log(middleNumber); // 3
```

السبب في استخدام `Math.floor` هو أن نتيجة قسمة عدد فردي على `2` ستكون عددًا عشريًا. تقوم `Math.floor` بتقريب العدد إلى أقرب عدد صحيح أقل.

أعلن عن متغير `oddListMedian` وقم بتعيين نتيجة إيجاد الرقم الأوسط من `testArr1` إليه. ثم سجل المتغير `oddListMedian` في وحدة التحكم.

افتح وحدة التحكم لرؤية النتيجة.

# --hints--

يجب أن يكون لديك متغير `oddListMedian`.

```js
assert.isDefined(oddListMedian);
```

يجب أن تعيّن نتيجة إيجاد الرقم الأوسط لـ `testArr1` إلى المتغير `oddListMedian`. تذكر الرجوع إلى المثال السابق: `arr[Math.floor(arr.length / 2)];`

```js
assert.strictEqual(oddListMedian, 3);
```

يجب أن يكون لديك بيان في وحدة التحكم يسجل المتغير `oddListMedian`.

```js
assert.match(code, /console\.log\(\s*oddListMedian\s*\);?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const testArr1 = [1, 2, 3, 4, 5];
const testArr2 = [1, 2, 3, 4, 5, 6];
const isEven = testArr2.length % 2 === 0;
console.log(isEven);


--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
