---
id: 687d8e05a529a4838f1c961c
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Em seguida, remova a palavra-chave `pass` dentro da classe `Node`, depois dê à classe `Node` um método `__init__` com dois parâmetros, `self` e `element`.

Dentro do método, atribua `element` a `self.element`. Isso armazena o valor passado ao criar um nó.

# --hints--

Você deve definir um método `__init__` dentro da sua classe `Node`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Seu método `__init__` deve ter dois parâmetros, `self` e `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Você deve atribuir `element` a `self.element` dentro do método `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
