---
id: 687d7deea6a2a22fa8816a73
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Nas lições anteriores, você aprendeu sobre algumas estruturas de dados comuns. Neste workshop, você implementará uma linked list.

Primeiro, crie uma classe chamada `LinkedList` para representar a estrutura que conterá os nós da sua lista. Adicione a palavra-chave `pass` dentro da classe para que o teste possa passar.

# --hints--

Você deve criar uma classe chamada `LinkedList`. Lembre-se de adicionar a palavra-chave `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
