---
id: 687d8c6f066877768f7ee441
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro da sua classe `LinkedList`, substitua a palavra-chave `pass` por uma classe interna chamada `Node`. Isso representará nós individuais que armazenam dados na lista.

# --hints--

Você deve criar uma classe `Node` dentro da classe `LinkedList`. Lembre-se de adicionar a palavra-chave `pass` dentro da classe `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
