---
id: 687e366b3004a9831e4cbd41
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Adicione um método chamado `is_empty` que retorna `self.length == 0`.

Você usará isso para acompanhar quantos nós estão na lista. Então, com `self.length == 0`, ele retorna `True`, significando que a lista não tem nós. 

Se `self.length` for maior que `0`, ele retorna `False`, significando que a lista tem pelo menos um nó.

# --hints--

Você deve criar um método `is_empty` que retorne `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
