---
id: 687e344cc1993f7316af030c
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro da classe `LinkedList`, adicione um método `__init__` que defina `self.length` como `0`. Isso inicializa um contador para rastrear quantos nós (elementos) estão atualmente na lista.

# --hints--

Você deve definir `self.length` como `0` dentro do método `__init__` da classe `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
