---
id: 68800a64b06e80d9a31c131b
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Ainda no método `__init__` da classe `LinkedList`, defina `self.head` como `None`. Isso atua como uma referência para o primeiro nó da lista. A lista está vazia no momento, então você precisa defini-la como `None`.


# --hints--

Você deve definir `self.head` como `None` dentro do método `__init__` da classe `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
