---
id: 688003f589b875af60273dd9
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Ainda dentro do método `__init__`, atribua `None` a `self.next`. Isso define o ponteiro `next` como `None`, significando que o novo nó ainda não aponta para nenhum outro nó na lista.

Depois, `next` pode ser atualizado para apontar para outro nó para formar a cadeia de uma lista encadeada.

# --hints--

Dentro do método `__init__`, você deve inicializar o atributo `next` com `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
