---
id: 68857e0a49baf4751f737a2a
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Agora, depois do método `is_empty`, crie um método chamado `add` dentro da classe `LinkedList`. O novo método deve ter dois parâmetros, `self` e `element`.

Dentro do seu novo método, crie uma variável chamada `node` e atribua a ela uma nova instância de `Node` usando `self.Node(element)`. Isso criará um novo nó para ser adicionado à lista encadeada.

# --hints--

Você deve criar um método `add` dentro da classe `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Seu método `add` deve ter dois parâmetros, `self` e `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Você deve atribuir `self.Node(element)` a `node` no método `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
