---
id: 68858a1b4e867acd87f6f7e8
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Se a lista estiver vazia, então o novo nó se torna o primeiro elemento (o `head`). Dentro do método `add`, crie uma declaração `if` que verifica se a lista está vazia usando `self.is_empty()` e atribua `node` a `self.head`.

# --hints--

Você deve verificar se a lista está vazia usando o método `is_empty()`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_conditions()[0]
  assert _cond.is_equivalent("self.is_empty()")
`)})
```

Você deve atribuir `node` a `self.head` dentro da declaração `if`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_body().is_equivalent("self.head = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
