---
id: 688649c8673f842f49228a5e
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Se a lista não estiver vazia, o método deve encontrar o último nó e ligar o novo nó ao seu ponteiro `next`.

Para fazer isso, crie um bloco `else` após o bloco `if`. Dentro do bloco `else`, crie uma variável `current_node` e atribua `self.head` a ela.

# --hints--

Você deve atribuir `self.head` a `current_node` dentro do bloco `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
