---
id: 68867c3898b09115d8c5140f
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

O loop `while` que você acabou de criar permite que você percorra cada nó na lista até que o último nó seja alcançado. Quando isso acontecer, você pode adicionar o nó à lista.

No bloco `else`, logo após o bloco `while`, atribua `node` a `current_node.next`.

# --hints--

Você deve atribuir `node` a `current_node.next` no bloco `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
