---
id: 6888f9827637322d4f309192
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

Dentro do método `remove`, crie uma variável chamada `current_node` e atribua a ela o valor de `self.head`.

Isto inicia a travessia da lista ligada a partir do primeiro nó. `current_node` será usado para acompanhar o nó que está sendo examinado enquanto você procura o nó que contém o elemento especificado para remover.

# --hints--

Você deve atribuir `self.head` a `current_node` dentro do método `remove`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").has_stmt("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
