---
id: 68b84df2648b4c96b25022f3
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Vamos testar sua lista ligada até agora! Na parte inferior do seu código, crie uma instância de `LinkedList` chamada `my_list`. Então, usando o método `is_empty` que você criou na etapa anterior, imprima `my_list.is_empty()`.

# --hints--

Você deve criar uma instância de `LinkedList` chamada `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Você deve imprimir `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
