---
id: 68b9e951e4892721632acc26
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

Ótimo! Agora vamos testar a adição de elementos à sua lista encadeada. No final do seu código, depois de criar uma instância de `LinkedList` chamada `my_list` e imprimir `my_list.is_empty()`, adicione `1` e `2` à lista usando o método `add`.

Então, imprima `my_list.is_empty()` novamente para confirmar que a lista não está mais vazia. Finalmente, imprima `my_list.length` para ver o comprimento atual da lista encadeada.

# --hints--

Você deve adicionar `1` à lista usando o método `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Você deve adicionar `2` à lista usando o método `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Você deve imprimir `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Você deve imprimir `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
