---
id: 66db529d37ad966480ebb633
title: بناء برنامج اختيار الغداء
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

في هذا المختبر، ستبني برنامجًا يساعد في إدارة خيارات الغداء. ستتعامل مع مصفوفة من وجبات الغداء، تضيف وتحذف عناصر من المصفوفة، وتختار خيار غداء عشوائيًا.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب عليك إنشاء متغير `lunches` وتعيينه إلى مصفوفة فارغة تُستخدم لتخزين عناصر الغداء.

2. يجب عليك إنشاء دالة `addLunchToEnd` تأخذ مصفوفة كمعلمة أولى وسلسلة نصية كمعلمة ثانية. يجب أن تقوم الدالة بما يلي:
    - إضافة السلسلة النصية إلى نهاية المصفوفة.
    - تسجيل السلسلة النصية `[Lunch Item] added to the end of the lunch menu.` في وحدة التحكم، حيث `[Lunch Item]` هي السلسلة النصية الممررة إلى الدالة.
    - إرجاع المصفوفة المحدثة.

3. يجب عليك إنشاء دالة `addLunchToStart` تأخذ مصفوفة كمعلمة أولى وسلسلة نصية كمعلمة ثانية. يجب أن تقوم الدالة بما يلي:
    - إضافة السلسلة النصية إلى بداية المصفوفة.
    - تسجيل السلسلة النصية `[Lunch Item] added to the start of the lunch menu.` في وحدة التحكم، حيث `[Lunch Item]` هي السلسلة النصية الممررة إلى الدالة.
    - إرجاع المصفوفة المحدثة.

4. يجب عليك إنشاء دالة `removeLastLunch` تأخذ مصفوفة كمعامل. يجب أن تقوم الدالة بما يلي:
   - إزالة العنصر الأخير من المصفوفة.
   - إذا نجحت عملية الإزالة، سجّل في وحدة التحكم السلسلة النصية `[Lunch Item] removed from the end of the lunch menu.` حيث يكون `[Lunch Item]` هو العنصر الذي أُزيل من المصفوفة.
   - إذا كانت المصفوفة فارغة، سجّل في وحدة التحكم السلسلة النصية `"No lunches to remove."`.
   - إرجاع المصفوفة المحدثة.

5. يجب عليك إنشاء دالة `removeFirstLunch` تأخذ مصفوفة كمعامل. يجب أن تقوم الدالة بما يلي:
   - إزالة العنصر الأول من المصفوفة.
   - إذا نجحت عملية الإزالة، سجّل في وحدة التحكم السلسلة النصية `[Lunch Item] removed from the start of the lunch menu.` حيث يمثل `[Lunch Item]` العنصر الذي أُزيل من المصفوفة.
   - إذا كانت المصفوفة فارغة، سجّل في وحدة التحكم السلسلة النصية `"No lunches to remove."`.
   - أرجع المصفوفة المحدثة.

6. يجب عليك إنشاء دالة `getRandomLunch` تأخذ مصفوفة كمعطى. يجب أن تقوم الدالة بما يلي:
   - اختيار عنصر عشوائي من المصفوفة.
   - إذا تم الاختيار بنجاح، سجل في وحدة التحكم السلسلة النصية `Randomly selected lunch: [Lunch Item]`، حيث يمثل `[Lunch Item]` عنصرًا عشوائيًا من المصفوفة.
   - إذا كانت المصفوفة فارغة، سجل في وحدة التحكم السلسلة النصية `"No lunches available."`.

7. يجب أن تنشئ دالة `showLunchMenu` تأخذ مصفوفة كمعامل لها و:
   - إذا كانت هناك عناصر في المصفوفة، تسجل السلسلة النصية `Menu items: [Lunch Item], [Lunch Item]...` في وحدة التحكم، حيث كل `[Lunch item]` هو أحد العناصر في المصفوفة، بالترتيب.
   - إذا كانت المصفوفة فارغة، تسجل السلسلة النصية `"The menu is empty."` في وحدة التحكم.

# --hints--

يجب إعلان متغير `lunches` وتعيينه إلى مصفوفة فارغة لتخزين عناصر الغداء.

```js
assert.isArray(lunches);
```

يجب تعريف دالة `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

يجب أن تحتوي الدالة `addLunchToEnd` على معلمتين.

```js
assert.lengthOf(addLunchToEnd, 2);
```

يجب أن تُسجل الدالة `addLunchToEnd(lunches, "Tacos")` السلسلة النصية `"Tacos added to the end of the lunch menu."` في وحدة التحكم.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` يجب أن تُرجع `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

يجب تعريف دالة `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

يجب أن تحتوي الدالة `addLunchToStart` على معلمتين.

```js
assert.lengthOf(addLunchToStart, 2);
```

يجب أن تُسجل الدالة `addLunchToStart(lunches, "Sushi")` السلسلة النصية `"Sushi added to the start of the lunch menu."` في وحدة التحكم.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` يجب أن تُرجع `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

يجب تعريف دالة `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

يجب أن تحتوي الدالة `removeLastLunch` على معلمة واحدة.

```js
assert.lengthOf(removeLastLunch, 1);
```

عندما تكون المصفوفة المُدخلة فارغة، يجب أن تُسجّل الدالة `removeLastLunch` السلسلة النصية `"No lunches to remove."` في وحدة التحكم.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

يجب أن تسجل الدالة `removeLastLunch(["Stew", "Soup", "Toast"])` السلسلة النصية `"Toast removed from the end of the lunch menu."` في وحدة التحكم.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` يجب أن تُرجع `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

يجب تعريف دالة `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

يجب أن تحتوي الدالة `removeFirstLunch` على معلمة واحدة.

```js
assert.lengthOf(removeFirstLunch, 1);
```

عندما تكون المصفوفة المُدخلة فارغة، يجب أن تسجل الدالة `removeFirstLunch` السلسلة النصية `"No lunches to remove."` في وحدة التحكم.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

يجب أن تسجل الدالة `removeFirstLunch(["Salad", "Eggs", "Cheese"])` السلسلة النصية `"Salad removed from the start of the lunch menu."` في وحدة التحكم.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` يجب أن تُرجع `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

يجب أن تُرجع الدوال `addLunchToEnd`، `addLunchToStart`، `removeLastLunch`، و`removeFirstLunch` نفس المصفوفة التي تم تمريرها كمعطى بعد تحديثها.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

يجب تعريف دالة `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

يجب أن تحتوي الدالة `getRandomLunch` على معلمة واحدة.

```js
assert.lengthOf(getRandomLunch, 1);
```

عندما تكون المصفوفة المُدخلة فارغة، يجب أن تسجل الدالة `getRandomLunch` السلسلة النصية `"No lunches available."` في وحدة التحكم.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

عندما لا تكون المصفوفة المدخلة فارغة، يجب أن تسجل الدالة `getRandomLunch` سلسلة نصية بصيغة `Randomly selected lunch: [Lunch Item]` في وحدة التحكم.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

يجب ألا تغير الدالة `getRandomLunch` المصفوفة التي تم تمريرها إليها كمعطى.

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

يجب تعريف دالة `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

يجب أن تحتوي الدالة `showLunchMenu` على معلمة واحدة.

```js
assert.lengthOf(showLunchMenu, 1);
```

عندما تكون المصفوفة المُدخلة فارغة، يجب أن تسجل الدالة `showLunchMenu` السلسلة النصية `"The menu is empty."` في وحدة التحكم.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

يجب أن تُسجل الدالة `showLunchMenu(["Greens", "Corns", "Beans"])` النص `"Menu items: Greens, Corns, Beans"` في وحدة التحكم.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

يجب أن تسجل الدالة `showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` السلسلة `"Menu items: Pizza, Burger, Fries, Salad"` في وحدة التحكم.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
