---
id: 587d7b88367417b2b2512b44
title: 파라미터와 함께 화살표 함수 작성하기
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

Just like a regular function, you can pass arguments into an arrow function.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` 은 `8`을 반환할 것입니다.

화살표 함수 파라미터가 단 하나인 경우, 파라미터를 감싸는 괄호는 생략할 수 있습니다.

```js
const doubler = item => item * 2;
```

화살표 함수도 두 개 이상의 인수를 전달할 수 있습니다.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` 는 `8`을 반환할 것입니다.

# --instructions--

`arr2`의 값을 `arr1`로 추가하는 함수 `myConcat`을 화살표 함수 문법으로 재작성하세요.

# --hints--

키워드 `var`는 다른 선언 키워드로 교체해야 합니다.

```js
assert.notMatch(code, /var/g);
```

`myConcat`은 (`const`를 이용하여) 상수로 선언해야 합니다.

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat`은 두 개의 파라미터를 가진 화살표 함수여야 합니다.

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()`은 `[1, 2, 3, 4, 5]`을 반환해야 합니다.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

키워드 `function`은 사용하지 않습니다.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
