---
id: 587d7b8c367417b2b2512b58
title: export default로 대체 수단 내보내기 만들기
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

`export` 수업에서 <dfn>이름 있는 내보내기</dfn>라고 하는 구문에 대해 배웠습니다. 이를 통해 여러 함수와 변수를 다른 파일에서 사용할 수 있게 만들었습니다.

알아야 할 또 다른 `export` 구문이 있는데, 이것은 <dfn>export default</dfn>로 알려져 있습니다. 보통 파일에서 하나의 값만 내보낼 때 이 구문을 사용합니다. 또한 파일이나 모듈의 대체 수단 값을 만들 때도 사용합니다.

아래는 `export default`를 사용한 예시입니다:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

첫 번째는 이름 있는 함수이고, 두 번째는 익명 함수입니다.

`export default`는 모듈이나 파일의 대체 수단 값을 선언하는 데 사용되므로, 각 모듈이나 파일에는 하나의 기본 내보내기 값만 있을 수 있습니다. 또한 `export default`는 `var`, `let`, 또는 `const`와 함께 사용할 수 없습니다.

# --instructions--

다음 함수가 모듈의 대체 수단 값이 되어야 합니다. 그렇게 하기 위해 필요한 코드를 추가하세요.

# --hints--

코드는 `export` 대체 수단을 사용해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
