---
id: 587d7b8a367417b2b2512b4d
title: 함수 매개변수로 객체를 전달할 때 구조 분해 할당 사용하기
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

경우에 따라 함수 인수 자체에서 객체를 구조 분해할 수 있습니다.

아래 코드를 살펴보세요:

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

이 코드는 함수에 전달된 객체를 효과적으로 구조 분해합니다. 이 작업은 다음과 같이 인라인으로도 할 수 있습니다:

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

위 함수에 `profileData`가 전달되면, 함수 매개변수에서 값이 구조 분해되어 함수 내에서 사용됩니다.

# --instructions--

함수 `half`의 인수 내에서 구조 분해 할당을 사용하여 함수 내부에 `max`와 `min`만 전달하세요.

# --hints--

`stats`는 `object`여야 합니다.

```js
assert(typeof stats === 'object');
```

`half(stats)`는 `28.015`여야 합니다.

```js
assert(half(stats) === 28.015);
```

구조 분해를 사용하세요.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

구조 분해된 매개변수를 사용하세요.

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
