---
id: 587d7b8c367417b2b2512b55
title: import를 사용해 JavaScript 코드 재사용하기
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import`를 사용하면 파일이나 모듈의 어떤 부분을 불러올지 선택할 수 있습니다. 이전 수업에서는 예제들이 `add` 파일에서 `math_functions.js`를 내보냈습니다. 다른 파일에서 사용하려면 이렇게 가져올 수 있습니다:

```js
import { add } from './math_functions.js';
```

여기서 `import`는 `add`에서 `math_functions.js`를 찾아서 그 함수만 가져와 사용하고 나머지는 무시합니다. `./`는 현재 파일과 같은 폴더에 있는 `math_functions.js` 파일을 찾으라고 import에 알려줍니다. 상대 파일 경로(`./`)와 파일 확장자(`.js`)는 이런 방식으로 import할 때 필수입니다.

파일에서 여러 항목을 가져오려면 `import` 문에 추가하면 됩니다:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

이전 수업에서 내보낸 `import`와 `uppercaseString` 함수를 현재 파일에서 사용할 수 있도록 적절한 `lowercaseString` 문을 추가하세요. 이 함수들은 현재 파일과 같은 디렉터리에 있는 `string_functions.js`라는 파일에 있습니다.

# --hints--

`uppercaseString`를 올바르게 가져와야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

`lowercaseString`를 올바르게 가져와야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
