---
id: 587d7b8a367417b2b2512b4e
title: 템플릿 리터럴을 사용해 문자열 만들기
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

ES6의 새로운 기능 중 하나는 <dfn>템플릿 리터럴</dfn>입니다. 이것은 복잡한 문자열을 더 쉽게 만들 수 있게 해주는 특별한 유형의 문자열입니다.

템플릿 리터럴을 사용하면 여러 줄 문자열을 만들고 문자열 보간 기능을 활용해 문자열을 생성할 수 있습니다.

아래 코드를 살펴보세요:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

콘솔에는 `Hello, my name is Zodiac Hasbro!`와 `I am 56 years old.` 문자열이 출력됩니다.

여기서 많은 일이 일어났습니다. 우선, 예제는 백틱(`` ` ``), not quotes (`'` or `"`), to wrap the string. Secondly, notice that the string is multi-line, both in the code and the output. This saves inserting `\n` within strings. The `${variable}` syntax used above is a placeholder. Basically, you won't have to use concatenation with the `+` operator anymore. To add variables to strings, you just drop the variable in a template string and wrap it with `${` and `}`. Similarly, you can include other expressions in your string literal, for example `${a + b}`)을 사용했습니다. 이 새로운 문자열 생성 방식은 더 견고한 문자열을 만들 수 있는 유연성을 제공합니다.

# --instructions--

템플릿 리터럴 구문을 백틱과 함께 사용해 리스트 요소(`li`) 문자열 배열을 만드세요. 각 리스트 요소의 텍스트는 `failure` 객체의 `result` 속성 배열 요소 중 하나여야 하며, `class` 속성에 `text-warning` 값을 가져야 합니다. `makeList` 함수는 리스트 아이템 문자열 배열을 반환해야 합니다.

원하는 출력(아래에 표시됨)을 얻기 위해 반복자 메서드(어떤 종류의 반복문도 가능)를 사용하세요.

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList`는 `result failure` 메시지를 포함하는 배열이어야 합니다.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList`는 지정된 출력과 같아야 합니다.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

템플릿 문자열과 표현식 보간을 사용해야 합니다.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

반복자를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
