---
id: 587d7b88367417b2b2512b46
title: 함수의 기본 매개변수 설정하기
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

더 유연한 함수를 만들 수 있도록 ES6에서는 함수에 대한 <dfn>기본 매개변수</dfn>를 도입했습니다.

다음 코드를 확인하세요:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

콘솔에는 문자열 `Hello John`와 `Hello Anonymous`이 출력됩니다.

기본 매개변수는 인수가 지정되지 않았을 때(즉, undefined일 때) 작동합니다. 위 예시에서 보듯이, 매개변수 `name`에 값을 제공하지 않으면 기본값 `Anonymous`이 할당됩니다. 원하는 만큼 많은 매개변수에 기본값을 추가할 수 있습니다.

# --instructions--

함수 `increment`를 수정하여 `number`이 지정되지 않았을 경우 `value`에 1을 더하도록 기본 매개변수를 추가하세요.

# --hints--

`increment(5, 2)`의 결과는 `7`이어야 합니다.

```js
assert(increment(5, 2) === 7);
```

`increment(5)`의 결과는 `6`이어야 합니다.

```js
assert(increment(5) === 6);
```

`1`에 대해 기본 매개변수 값 `value`을 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
