---
id: 587d7b87367417b2b2512b42
title: const로 선언한 배열 변경하기
challengeType: 1
forumTopicId: 301206
dashedName: mutate-an-array-declared-with-const
---

# --description--

`const`에 익숙하지 않다면 <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/declare-a-read-only-variable-with-the-const-keyword" target="_blank" rel="noopener noreferrer nofollow">`const` 키워드에 관한 이 도전 과제</a>를 확인하세요.

`const` 선언은 현대 JavaScript에서 다양한 용도로 사용됩니다.

일부 개발자는 값을 다시 할당해야 할 필요가 있을 때를 제외하고는 기본적으로 모든 변수를 `const`로 할당하는 것을 선호합니다. 그런 경우에만 `let`를 사용합니다.

하지만 `const`를 사용해 변수에 할당된 객체(배열과 함수 포함)는 여전히 변경 가능하다는 점을 이해하는 것이 중요합니다. `const` 선언은 변수 식별자의 재할당만 막아줍니다.

```js
const s = [5, 6, 7];
s = [1, 2, 3];
s[2] = 45;
console.log(s);
```

`s = [1, 2, 3]`는 오류를 발생시킵니다. 해당 줄을 주석 처리한 후 `console.log`는 값 `[5, 6, 45]`를 출력합니다.

보시다시피, 객체 `[5, 6, 7]` 자체를 변경할 수 있고 변수 `s`는 변경된 배열 `[5, 6, 45]`를 계속 가리킵니다. 모든 배열과 마찬가지로 `s`의 배열 요소는 변경 가능하지만 `const`가 사용되었기 때문에 변수 식별자 `s`를 사용해 다른 배열을 가리키도록 대입 연산자를 사용할 수 없습니다.

# --instructions--

배열이 `const s = [5, 7, 2]`로 선언되어 있습니다. 다양한 요소 할당을 사용해 배열을 `[2, 5, 7]`로 변경하세요.

# --hints--

`const` 키워드를 바꾸면 안 됩니다.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`s`는 상수 변수여야 합니다(`const` 사용).

```js
assert(__helpers.removeJSComments(code).match(/const\s+s/g));
```

원래 배열 선언을 변경하면 안 됩니다.

```js
assert(__helpers.removeJSComments(code).match(
/const\s+s\s*=\s*\[\s*5\s*,\s*7\s*,\s*2\s*\]\s*;?/g
));
```

`s`는 `[2, 5, 7]`와 같아야 합니다.

```js
assert.deepEqual(s, [2, 5, 7]);
```

# --seed--

## --seed-contents--

```js
const s = [5, 7, 2];
function editInPlace() {
  // Only change code below this line

  // Using s = [2, 5, 7] would be invalid

  // Only change code above this line
}
editInPlace();
```

# --solutions--

```js
const s = [5, 7, 2];
function editInPlace() {
  s[0] = 2;
  s[1] = 5;
  s[2] = 7;
}
editInPlace();
```
