---
id: 587d7b87367417b2b2512b43
title: 간결한 익명 함수를 작성하기 위해 화살표 함수 사용하기
challengeType: 1
forumTopicId: 301211
dashedName: use-arrow-functions-to-write-concise-anonymous-functions
---

# --description--

JavaScript에서는 특히 함수를 다른 함수의 인수로 전달할 때 함수에 이름을 붙일 필요가 없는 경우가 많습니다. 대신 인라인 함수를 만듭니다. 이러한 함수는 다른 곳에서 재사용하지 않기 때문에 이름을 붙일 필요가 없습니다.

이를 위해 보통 다음 구문을 사용합니다:

```js
const myFunc = function() {
  const myVar = "value";
  return myVar;
}
```

ES6는 익명 함수를 이렇게 작성하지 않아도 되도록 구문 설탕을 제공합니다. 대신 **화살표 함수 구문**을 사용할 수 있습니다:

```js
const myFunc = () => {
  const myVar = "value";
  return myVar;
}
```

함수 본문이 없고 반환값만 있을 때, 화살표 함수 구문은 키워드 `return`와 코드를 감싸는 괄호를 생략할 수 있게 해줍니다. 이는 작은 함수를 한 줄 문장으로 간소화하는 데 도움이 됩니다:

```js
const myFunc = () => "value";
```

이 코드는 기본적으로 문자열 `value`을 반환합니다.

# --instructions--

`magic` 변수에 할당된, `new Date()`을 반환하는 함수를 화살표 함수 구문을 사용하도록 다시 작성하세요. 또한 `var` 키워드를 사용해 정의된 것이 없도록 하세요.

# --hints--

`var` 키워드를 교체해야 합니다.

```js
assert.notMatch(code, /var/g)
```

`magic`는 상수 변수여야 합니다 (`const`을 사용하여).

```js
assert.match(code, /const\s+magic/g)
```

`magic`는 `function`여야 합니다.

```js
assert(typeof magic === 'function');
```

`magic()`은 올바른 날짜를 반환해야 합니다.

```js
assert(magic().setHours(0, 0, 0, 0) === new Date().setHours(0, 0, 0, 0));
```

`function` 키워드는 사용하지 않아야 합니다.

```js
assert.notMatch(code, /function/g)
```

# --seed--

## --seed-contents--

```js
var magic = function() {
  return new Date();
};
```

# --solutions--

```js
const magic = () => {
  return new Date();
};
```
