---
id: 587d7b89367417b2b2512b4a
title: 중첩 객체에서 변수를 할당하기 위해 구조 분해 할당 사용하기
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

이전 두 수업의 원칙을 사용해 중첩 객체에서 값을 구조 분해할 수 있습니다.

이전 예제와 비슷한 객체를 사용하면:

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

객체 프로퍼티의 값을 추출해 같은 이름의 변수에 할당하는 방법은 다음과 같습니다:

```js
const { johnDoe: { age, email }} = user;
```

그리고 객체 프로퍼티의 값을 다른 이름의 변수에 할당하는 방법은 다음과 같습니다:

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

두 할당문을 동등한 구조 분해 할당으로 바꾸세요. `lowToday` 객체에서 `highToday`와 `today.low`의 값을 `today.high`와 `LOCAL_FORECAST` 변수에 할당해야 합니다.

# --hints--

ES5 할당 구문을 제거하세요.

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

구조 분해를 사용해 `lowToday` 변수를 만드세요.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

구조 분해를 사용해 `highToday` 변수를 만드세요.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday`는 `64`와 같아야 하고 `highToday`는 `77`와 같아야 합니다.

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
