---
id: 587d7b8c367417b2b2512b54
title: 객체 접근 제어를 위한 getter와 setter 사용하기
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

객체에서 값을 얻거나 객체 내 속성의 값을 설정할 수 있습니다.

이것을 전통적으로 <dfn>getter</dfn>와 <dfn>setter</dfn>라고 부릅니다.

Getter 함수는 사용자가 객체의 비공개 변수를 직접 접근하지 않고도 그 값을 단순히 반환(얻기)하도록 설계되었습니다.

Setter 함수는 전달받은 값을 바탕으로 객체의 비공개 변수 값을 수정(설정)하도록 설계되었습니다. 이 변경은 계산을 포함할 수도 있고 이전 값을 완전히 덮어쓸 수도 있습니다.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

콘솔에는 `anonymous`와 `newAuthor` 문자열이 표시됩니다.

getter와 setter를 호출하는 구문에 주목하세요. 함수처럼 보이지도 않습니다. Getter와 setter는 내부 구현 세부사항을 숨기기 때문에 중요합니다.

**참고:** 비공개 변수 이름 앞에 밑줄(`_`)을 붙이는 것이 관례입니다. 하지만 이 관례 자체가 변수를 비공개로 만드는 것은 아닙니다.

# --instructions--

`class` 키워드를 사용해 `Thermostat` 클래스를 만드세요. `constructor`는 화씨 온도를 받습니다.

클래스 내에 섭씨 온도를 얻는 `getter`와 섭씨 온도를 받는 `setter`를 만드세요.

`C = 5/9 * (F - 32)`와 `F = C * 9.0 / 5 + 32`를 기억하세요. 여기서 `F`은 화씨 온도 값이고, `C`은 같은 온도의 섭씨 값입니다.

**참고:** 이 구현에서는 클래스 내부에서 온도를 화씨나 섭씨 중 한 가지 단위로 추적합니다.

이것이 getter와 setter의 힘입니다. 다른 사용자를 위한 API를 만들고, 어떤 단위를 추적하든 올바른 결과를 얻을 수 있습니다.

즉, 사용자가 구현 세부사항을 알 필요 없도록 추상화하는 것입니다.

# --hints--

`Thermostat`는 `class`이며, `constructor` 메서드를 정의해야 합니다.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

`class` 키워드를 사용해야 합니다.

```js
assert.match(code, /class/);
```

`Thermostat`는 인스턴스화할 수 있어야 합니다.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

화씨 값을 인자로 인스턴스화할 때 `Thermostat`가 올바른 `temperature`를 설정해야 합니다.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

`getter`가 정의되어야 합니다.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

`setter`가 정의되어야 합니다.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

섭씨 값을 인자로 `setter`를 호출하면 `temperature`가 설정되어야 합니다.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
