---
id: 587d7b87367417b2b2512b40
title: var와 let 키워드의 범위 비교
challengeType: 1
forumTopicId: 301195
dashedName: compare-scopes-of-the-var-and-let-keywords
---

# --description--

`let`에 익숙하지 않다면 <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/explore-differences-between-the-var-and-let-keywords" target="_blank" rel="noopener noreferrer nofollow"><code>let</code>과 <code>var</code>의 차이에 관한 이 도전 과제</a>를 확인하세요.

`var` 키워드로 변수를 선언하면 전역으로 선언되거나 함수 내부에서 선언하면 지역으로 선언됩니다.

`let` 키워드는 비슷하게 동작하지만 몇 가지 추가 기능이 있습니다. `let` 키워드로 블록, 문장, 표현식 내부에서 변수를 선언하면 그 범위는 해당 블록, 문장, 표현식으로 제한됩니다.

예를 들어:

```js
var numArray = [];
for (var i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

여기서 콘솔은 `[0, 1, 2]`와 `3` 값을 출력합니다.

`var` 키워드로 `i`는 전역으로 선언됩니다. 그래서 `i++`가 실행될 때 전역 변수를 업데이트합니다. 이 코드는 다음과 유사합니다:

```js
var numArray = [];
var i;
for (i = 0; i < 3; i++) {
  numArray.push(i);
}
console.log(numArray);
console.log(i);
```

여기서 콘솔은 `[0, 1, 2]`와 `3` 값을 출력합니다.

이 동작은 `for` 루프 안에서 `i` 변수를 사용하는 함수를 만들고 나중에 사용하기 위해 저장하면 문제가 됩니다. 저장된 함수는 항상 업데이트된 전역 `i` 변수의 값을 참조하기 때문입니다.

```js
var printNumTwo;
for (var i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
```

여기서 콘솔은 `3` 값을 출력합니다.

보시다시피 `printNumTwo()`는 3을 출력하며 2가 아닙니다. 이는 `i`에 할당된 값이 업데이트되었고 `printNumTwo()`가 함수가 for 루프에서 생성될 때의 `i` 값이 아니라 전역 `i`를 반환하기 때문입니다. `let` 키워드는 이러한 동작을 따르지 않습니다:

```js
let printNumTwo;
for (let i = 0; i < 3; i++) {
  if (i === 2) {
    printNumTwo = function() {
      return i;
    };
  }
}
console.log(printNumTwo());
console.log(i);
```

여기서 콘솔은 `2` 값을 출력하고 `i is not defined`라는 오류가 발생합니다.

`i`는 전역 범위에서 선언되지 않았기 때문에 정의되지 않았습니다. 이는 `for` 루프 문 내에서만 선언되었습니다. `printNumTwo()`는 올바른 값을 반환했는데, 이는 루프 문 내에서 `i` 키워드로 고유한 값(0, 1, 2)을 가진 세 개의 서로 다른 `let` 변수가 생성되었기 때문입니다.

# --instructions--

`i` 문에서 선언된 `if`가 함수 첫 줄에서 선언된 `i`와 별개의 변수이도록 코드를 수정하세요. 코드 어디에서도 `var` 키워드를 사용하지 않도록 주의하세요.

이 연습은 `var`와 `let` 키워드가 선언된 변수에 범위를 할당하는 방식의 차이를 보여주기 위해 설계되었습니다. 이 연습에서 사용한 것과 유사한 함수를 프로그래밍할 때 혼동을 피하려면 다른 변수 이름을 사용하는 것이 좋습니다.

# --hints--

코드에 `var`가 존재하면 안 됩니다.

```js
assert(!__helpers.removeJSComments(code).match(/var/g));
```

`i` 문에서 선언된 `if` 변수는 문자열 `block scope`와 같아야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/(i\s*=\s*).*\s*.*\s*.*\1('|")block\s*scope\2/g));
```

`checkScope()`는 문자열 `function scope`를 반환해야 합니다.

```js
assert(checkScope() === 'function scope');
```

# --seed--

## --seed-contents--

```js
function checkScope() {
  var i = 'function scope';
  if (true) {
    i = 'block scope';
    console.log('Block scope i is: ', i);
  }
  console.log('Function scope i is: ', i);
  return i;
}
```

# --solutions--

```js
function checkScope() {
  let i = 'function scope';
  if (true) {
    let i = 'block scope';
    console.log('Block scope i is: ', i);
  }
 
  console.log('Function scope i is: ', i);
  return i;
}
```
