---
id: 587d7b8c367417b2b2512b56
title: 코드 블록을 공유하기 위해 export 사용하기
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

수학 연산과 관련된 여러 함수가 들어 있는 `math_functions.js`라는 파일이 있다고 상상해 보세요. 그중 하나는 두 숫자를 받아서 합을 반환하는 `add`라는 변수에 저장되어 있습니다. 이 함수를 여러 JavaScript 파일에서 사용하려면 먼저 `export`해야 합니다.

```js
export const add = (x, y) => {
  return x + y;
}
```

위는 단일 함수를 export하는 일반적인 방법이지만, 다음과 같이 같은 결과를 얻을 수도 있습니다:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

변수나 함수를 export하면 다른 파일에서 가져와 코드를 다시 작성하지 않고 사용할 수 있습니다. export하려는 각 항목에 대해 첫 번째 예제를 반복하거나, 두 번째 예제의 export 문 안에 모두 넣어 여러 항목을 export할 수 있습니다:

```js
export { add, subtract };
```

# --instructions--

편집기에는 문자열 관련 함수 두 개가 있습니다. 원하는 방법으로 두 함수를 모두 export하세요.

# --hints--

`uppercaseString`를 올바르게 export해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

`lowercaseString`를 올바르게 export해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
