---
id: 587d7b89367417b2b2512b48
title: 전개 연산자를 사용해 배열을 제자리에서 평가하기
challengeType: 1
forumTopicId: 301222
dashedName: use-the-spread-operator-to-evaluate-arrays-in-place
---

# --description--

ES6는 <dfn>전개 연산자</dfn>를 도입하여, 여러 매개변수나 요소가 필요한 곳에서 배열과 다른 표현식을 확장할 수 있게 합니다.

아래 ES5 코드는 `apply()`를 사용해 배열에서 최대값을 계산합니다:

```js
var arr = [6, 89, 3, 45];
var maximus = Math.max.apply(null, arr);
```

`maximus`는 `89` 값을 갖게 됩니다.

`Math.max.apply(null, arr)`를 사용해야 했던 이유는 `Math.max(arr)`가 `NaN`를 반환하기 때문입니다. `Math.max()`는 쉼표로 구분된 인수를 기대하지만 배열은 아닙니다. 전개 연산자는 이 구문을 훨씬 읽기 쉽고 유지 관리하기 좋게 만듭니다.

```js
const arr = [6, 89, 3, 45];
const maximus = Math.max(...arr);
```

`maximus`는 `89` 값을 갖게 됩니다.

`...arr`는 풀린 배열을 반환합니다. 즉, 배열을 전개합니다. 하지만 전개 연산자는 함수의 인수나 배열 리터럴처럼 제자리에서만 작동합니다. 예를 들어:

```js
const spreaded = [...arr];
```

하지만 다음 코드는 작동하지 않습니다:

```js
const spreaded = ...arr;
```

# --instructions--

전개 연산자를 사용해 `arr1`의 모든 내용을 다른 배열 `arr2`에 복사하세요.

# --hints--

`arr2`는 `arr1`의 정확한 복사본이어야 합니다.

```js
assert(arr2.every((v, i) => v === arr1[i]) && arr2.length);
```

`...` 전개 연산자를 사용해 `arr1`를 복제해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/Array\(\s*\.\.\.arr1\s*\)|\[\s*\.\.\.arr1\s*\]/));
```

`arr2`는 `arr1`가 변경되어도 변하지 않아야 합니다.

```js
assert((arr1, arr2) => {
  arr1.push('JUN');
  return arr2.length < arr1.length;
});
```

# --seed--

## --seed-contents--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [];  // Change this line

console.log(arr2);
```

# --solutions--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [...arr1];
```
