---
id: 587d7b88367417b2b2512b44
title: 매개변수가 있는 화살표 함수 작성하기
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

일반 함수처럼 화살표 함수에도 인수를 전달할 수 있습니다.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)`는 `8` 값을 반환합니다.

화살표 함수에 매개변수가 하나만 있다면, 매개변수를 감싸는 소괄호를 생략할 수 있습니다.

```js
const doubler = item => item * 2;
```

화살표 함수에 여러 개의 인수를 전달하는 것도 가능합니다.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)`는 `8` 값을 반환합니다.

# --instructions--

`myConcat`의 내용을 `arr2`에 추가하는 `arr1` 함수를 화살표 함수 구문을 사용하도록 다시 작성하세요.

# --hints--

`var` 키워드를 교체해야 합니다.

```js
assert.notMatch(code, /var/g);
```

`myConcat`는 상수 변수여야 합니다 (`const` 사용).

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat`는 두 개의 매개변수를 가진 화살표 함수여야 합니다.

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()`는 `[1, 2, 3, 4, 5]`를 반환해야 합니다.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

`function` 키워드는 사용하지 않아야 합니다.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
