---
id: 587d7b8a367417b2b2512b4c
title: >-
  나머지 요소를 이용한 구조 분해
challengeType: 1
forumTopicId: 301218
dashedName: >-
  use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

배열 구조 분해를 할 때, 나머지 요소들을 별도의 배열로 모으고 싶을 때가 있습니다.

결과는 아래에 나온 `Array.prototype.slice()`와 비슷합니다:

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

콘솔에는 `1, 2`과 `[3, 4, 5, 7]` 값이 출력됩니다.

변수 `a`과 `b`는 배열에서 첫 번째와 두 번째 값을 가져옵니다. 그 후, 나머지 구문이 있기 때문에 `arr`는 나머지 값들을 배열 형태로 받습니다. 나머지 요소는 목록에서 마지막 변수로만 올바르게 작동합니다. 즉, 나머지 구문을 사용해 원래 배열의 마지막 요소를 제외한 하위 배열을 잡을 수는 없습니다.

# --instructions--

나머지 구문을 사용한 구조 분해 할당으로 `Array.prototype.slice()`의 동작을 흉내 내세요. `removeFirstTwo()`은 원래 배열 `list`에서 처음 두 요소를 제외한 하위 배열을 반환해야 합니다.

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])`은 `[3, 4, 5]`이어야 합니다.

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

`removeFirstTwo()`은 `list`를 수정하면 안 됩니다.

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

`Array.slice()`은 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

나머지 구문을 사용해야 합니다.

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
