---
id: 587d7b8a367417b2b2512b4f
title: 객체 속성 축약 구문을 사용해 간결한 객체 리터럴 선언 작성하기
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6는 객체 리터럴을 쉽게 정의할 수 있도록 멋진 지원을 추가했습니다.

다음 코드를 살펴보세요:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition`는 두 개의 속성을 포함하는 객체를 반환하는 간단한 함수입니다. ES6는 `x: x`을 반복해서 작성해야 하는 중복을 없애는 구문 설탕을 제공합니다. `x`를 한 번만 작성하면 내부적으로 `x: x`(또는 동등한 것)으로 변환됩니다. 위의 함수를 이 새로운 구문을 사용해 다시 작성한 예는 다음과 같습니다:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

객체 리터럴에서 객체 속성 축약 구문을 사용해 `name`, `age`, `gender` 속성을 가진 객체를 생성하고 반환하세요.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")`는 `{name: "Zodiac Hasbro", age: 56, gender: "male"}`을 반환해야 합니다.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

코드에서 `key:value`를 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
