---
id: 587d7b89367417b2b2512b4b
title: 배열에서 변수를 할당하기 위해 구조 분해 할당 사용하기
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6는 객체 구조 분해만큼 배열 구조 분해도 쉽게 만듭니다.

전개 연산자와 배열 구조 분해의 주요 차이점 중 하나는 전개 연산자가 배열의 모든 내용을 쉼표로 구분된 목록으로 풀어낸다는 점입니다. 따라서 변수에 할당할 요소를 선택할 수 없습니다.

배열 구조 분해를 사용하면 바로 그 작업을 할 수 있습니다:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

콘솔에는 `a`와 `b`의 값이 `1, 2`로 표시됩니다.

변수 `a`에는 배열의 첫 번째 값이 할당되고, `b`에는 배열의 두 번째 값이 할당됩니다. 또한 쉼표를 사용해 원하는 인덱스에 도달함으로써 구조 분해로 배열의 어떤 인덱스 값도 접근할 수 있습니다:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

콘솔에는 `a`, `b`, 그리고 `c`의 값이 `1, 2, 5`로 표시됩니다.

# --instructions--

구조 분해 할당을 사용해 `a`와 `b`의 값을 바꾸세요. 그렇게 해서 `a`이 `b`에 저장된 값을 받고, `b`이 `a`에 저장된 값을 받도록 하세요.

# --hints--

값 `a`는 교환 후 `6`이어야 합니다.

```js
assert(a === 6);
```

값 `b`는 교환 후 `8`이어야 합니다.

```js
assert(b === 8);
```

`a`와 `b`를 바꾸려면 배열 구조 분해를 사용하세요.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
