---
id: 587d7b8b367417b2b2512b53
title: 생성자 함수를 정의하기 위해 클래스 구문 사용하기
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6는 <dfn>클래스</dfn> 키워드를 사용하여 객체를 생성하는 새로운 구문을 제공합니다.

ES5에서는 `constructor` 함수를 정의하고 `new` 키워드를 사용해 객체를 인스턴스화하여 객체를 생성할 수 있습니다.

ES6에서는 `class` 선언에 `constructor` 메서드가 있으며, `new` 키워드로 호출됩니다. `constructor` 메서드가 명시적으로 정의되지 않으면 인수가 없는 상태로 암묵적으로 정의됩니다.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

`class` 키워드는 새로운 함수를 선언하며, 여기에 생성자가 추가된다는 점을 주목하세요. 이 생성자는 `new`가 호출되어 새 객체를 생성할 때 호출됩니다.

**참고:** ES6 클래스 이름에는 위 예시의 `SpaceShuttle`처럼 관례상 UpperCamelCase를 사용해야 합니다.

`constructor` 메서드는 클래스로 생성된 객체를 생성하고 초기화하는 특별한 메서드입니다. JavaScript 알고리즘과 자료 구조 인증의 객체 지향 프로그래밍 섹션에서 더 자세히 배우게 됩니다.

# --instructions--

`class` 키워드를 사용해 `constructor`을 작성하여 `Vegetable` 클래스를 만드세요.

`Vegetable` 클래스는 `name`로 전달되는 속성 `constructor`를 가진 vegetable 객체를 생성할 수 있게 합니다.

# --hints--

`Vegetable`는 `class`이며, `constructor` 메서드가 정의되어 있어야 합니다.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

`class` 키워드를 사용해야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable`는 인스턴스화할 수 있어야 합니다.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name`는 `carrot`를 반환해야 합니다.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
