---
id: 587d7b89367417b2b2512b49
title: 객체에서 변수를 할당할 때 구조 분해 할당 사용하기
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

구조 분해를 사용하면 값을 추출할 때 새 변수 이름을 할당할 수 있습니다. 값을 할당할 때 콜론 뒤에 새 이름을 적으면 됩니다.

지난 예제의 동일한 객체를 사용하면:

```js
const user = { name: 'John Doe', age: 34 };
```

다음과 같이 할당할 때 새 변수 이름을 줄 수 있습니다:

```js
const { name: userName, age: userAge } = user;
```

이를 "`user.name`의 값을 가져와 `userName`이라는 새 변수에 할당한다"라고 읽으면 됩니다. `userName`의 값은 문자열 `John Doe`이고, `userAge`의 값은 숫자 `34`입니다.

# --instructions--

두 개의 할당문을 동등한 구조 분해 할당으로 바꾸세요. `highToday` 객체에서 `highTomorrow`과 `today`의 값을 `tomorrow`과 `HIGH_TEMPERATURES` 변수에 할당해야 합니다.

# --hints--

ES5 할당 구문을 제거하세요.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

구조 분해를 사용해 `highToday` 변수를 만드세요.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

구조 분해를 사용해 `highTomorrow` 변수를 만드세요.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday`은 `77`와 같아야 하고 `highTomorrow`는 `80`와 같아야 합니다.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
