---
id: 587d7b8d367417b2b2512b59
title: 기본 내보내기 가져오기
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

지난 도전 과제에서 `export default`와 그 용도에 대해 배웠습니다. 기본 내보내기를 가져오려면 다른 `import` 구문을 사용해야 합니다. 다음 예제에서 `add`는 `math_functions.js` 파일의 기본 내보내기입니다. 가져오는 방법은 다음과 같습니다:

```js
import add from "./math_functions.js";
```

구문은 한 가지 중요한 부분에서 다릅니다. 가져온 값인 `add`는 중괄호(`{}`)로 둘러싸여 있지 않습니다. 여기서 `add`는 `math_functions.js` 파일의 기본 내보내기인 무엇이든지에 대한 변수 이름일 뿐입니다. 기본 내보내기를 가져올 때는 여기서 어떤 이름이든 사용할 수 있습니다.

# --instructions--

다음 코드에서 이 파일과 같은 디렉터리에 있는 `math_functions.js` 파일의 기본 내보내기를 가져오세요. 가져온 이름은 `subtract`로 지정하세요.

# --hints--

`subtract`를 `math_functions.js`에서 올바르게 가져와야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
