---
id: 587d7b88367417b2b2512b47
title: 함수 매개변수에 Rest 매개변수 사용하기
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

더 유연한 함수를 만들 수 있도록 ES6에서는 함수 매개변수에 <dfn>rest 매개변수</dfn>를 도입했습니다. rest 매개변수를 사용하면 가변 개수의 인수를 받는 함수를 만들 수 있습니다. 이 인수들은 배열에 저장되어 함수 내부에서 나중에 접근할 수 있습니다.

다음 코드를 확인하세요:

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

콘솔에는 문자열 `You have passed 3 arguments.`와 `You have passed 4 arguments.`이 출력됩니다.

rest 매개변수는 `arguments` 객체를 사용할 필요를 없애고, 함수에 전달된 매개변수 배열에 배열 메서드를 사용할 수 있게 해줍니다 `howMany`.

# --instructions--

함수 `sum`를 rest 매개변수를 사용하도록 수정하여, 함수 `sum`가 임의 개수의 인수를 받아 그 합을 반환할 수 있게 하세요.

# --hints--

`sum(0,1,2)`의 결과는 3이어야 합니다

```js
assert(sum(0, 1, 2) === 3);
```

`sum(1,2,3,4)`의 결과는 10이어야 합니다

```js
assert(sum(1, 2, 3, 4) === 10);
```

`sum(5)`의 결과는 5이어야 합니다

```js
assert(sum(5) === 5);
```

`sum()`의 결과는 0이어야 합니다

```js
assert(sum() === 0);
```

`sum`은 rest 매개변수 구문(`...`)을 `args` 매개변수에 사용하는 화살표 함수여야 합니다.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
