---
id: 587d7b8b367417b2b2512b50
title: ES6로 간결한 선언적 함수 작성하기
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

ES5에서 객체 안에 함수를 정의할 때는 다음과 같이 `function` 키워드를 사용해야 합니다:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

ES6에서는 객체 안에 함수를 정의할 때 `function` 키워드와 콜론을 완전히 제거할 수 있습니다. 이 구문 예시는 다음과 같습니다:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

객체 `setGear` 안의 함수 `bicycle`를 위에서 설명한 축약 구문을 사용하도록 리팩터링하세요.

# --hints--

기존의 함수 표현식은 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear`는 선언적 함수여야 합니다.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)`는 `gear` 값을 48로 변경해야 합니다.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
