---
id: 598f48a36c8c40764b4e52b3
title: 객체 변이 방지
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

이전 도전 과제에서 보았듯이 `const` 선언만으로는 데이터 변이를 완전히 막을 수 없습니다. 데이터가 변경되지 않도록 하려면 JavaScript가 제공하는 `Object.freeze` 함수를 사용해 변이를 방지할 수 있습니다.

객체를 변경하려는 모든 시도는 거부되며, 스크립트가 엄격 모드에서 실행 중이라면 오류가 발생합니다.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

`obj.review`와 `obj.newProp` 할당은 오류를 발생시키는데, 이는 기본적으로 편집기가 엄격 모드로 실행되기 때문이며, 콘솔에는 `{ name: "FreeCodeCamp", review: "Awesome" }` 값이 표시됩니다.

# --instructions--

이번 도전 과제에서는 `Object.freeze`를 사용해 수학 상수가 변경되지 않도록 막으세요. `MATH_CONSTANTS` 객체를 얼려서 `PI`의 값이 변경되거나 프로퍼티가 추가·삭제되지 않도록 해야 합니다.

# --hints--

`const` 키워드는 바꾸지 마세요.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS`는 상수 변수여야 합니다 (`const` 사용).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

`MATH_CONSTANTS`의 원래 선언은 변경하지 마세요.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI`는 `3.14`와 같아야 합니다.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
