---
id: 5cdafbe72913098997531682
title: catch로 거부된 프로미스 처리하기
challengeType: 1
forumTopicId: 301204
dashedName: handle-a-rejected-promise-with-catch
---

# --description--

`catch`는 프로미스가 거부되었을 때 사용하는 메서드입니다. 프로미스의 `reject` 메서드가 호출된 직후 실행됩니다. 구문은 다음과 같습니다:

```js
myPromise.catch(error => {
  
});
```

`error`는 `reject` 메서드에 전달되는 인수입니다.

# --instructions--

프로미스에 `catch` 메서드를 추가하세요. 콜백 함수의 매개변수로 `error`를 사용하고 `error`를 콘솔에 출력하세요.

# --hints--

프로미스에서 `catch` 메서드를 호출해야 합니다.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.catch\(/g)
);
```

`catch` 메서드는 `error`를 매개변수로 받는 콜백 함수를 가져야 합니다.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

`error`를 콘솔에 출력해야 합니다.

```js
assert(
  /\.catch\((function\(error\){|error|\(error\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.catch\(.*?error.*?console.log\(error\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to false to represent an unsuccessful response from a server
  let responseFromServer = false;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});

makeServerRequest.catch(error => {
  console.log(error);
});
```
