---
id: 5cdafbb0291309899753167f
title: JavaScript 프로미스 만들기
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

JavaScript에서 프로미스는 말 그대로 약속입니다 — 보통 비동기로 무언가를 하겠다는 약속을 하는 데 사용합니다. 작업이 완료되면 약속을 이행하거나 실패하게 됩니다. `Promise`는 생성자 함수이므로 `new` 키워드를 사용해 만들어야 합니다. 이 생성자는 두 개의 매개변수를 가진 함수를 인수로 받습니다 — `resolve`와 `reject`입니다. 이들은 프로미스의 결과를 결정하는 데 쓰이는 메서드입니다. 구문은 다음과 같습니다:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

`makeServerRequest`라는 새 프로미스를 만드세요. 생성자에 `resolve`와 `reject` 매개변수를 가진 함수를 전달하세요.

# --hints--

선언된 변수 `makeServerRequest`에 프로미스를 할당해야 합니다.

```js
assert(makeServerRequest instanceof Promise);
```

프로미스는 `resolve`와 `reject`를 매개변수로 받는 함수를 인수로 받아야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
