---
id: 5cdafbd72913098997531681
title: then으로 이행된 프로미스 처리하기
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

프로미스는 코드에서 시간이 얼마나 걸릴지 모르는 작업(즉, 비동기 작업), 보통 서버 요청이 있을 때 가장 유용합니다. 서버 요청을 하면 일정 시간이 걸리고, 완료된 후에는 보통 서버의 응답을 가지고 무언가를 하길 원합니다. 이는 `then` 메서드를 사용해서 할 수 있습니다.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

`then` 메서드는 프로미스가 결국 완료될 때 실행할 콜백 함수를 예약합니다 - 이행되거나 거부될 때입니다. `onFulfilled`와 `onRejected` 핸들러 중 하나가 현재 프로미스의 이행 또는 거부를 처리하기 위해 실행됩니다. 프로미스가 `resolve`로 이행되면 `onFulfilled` 핸들러가 호출됩니다.

```js
myPromise.then(result => {
  
});
```

`result`는 `resolve` 메서드에 전달된 인수에서 옵니다.

# --instructions--

프로미스에 `then` 메서드를 추가하세요. 콜백 함수의 매개변수로 `result`를 사용하고 `result`을 콘솔에 출력하세요.

# --hints--

프로미스에서 `then` 메서드를 호출해야 합니다.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

`then` 메서드는 `result`를 매개변수로 받는 콜백 함수를 가져야 합니다.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

`result`을 콘솔에 출력해야 합니다.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
