---
id: 5cdafbc32913098997531680
title: resolve와 reject로 프로미스 완료하기
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

프로미스는 세 가지 상태가 있습니다: `pending`, `fulfilled`, 그리고 `rejected`. 이전 도전 과제에서 만든 프로미스는 프로미스를 완료하는 방법을 추가하지 않아 영원히 `pending` 상태에 머물러 있습니다. 프로미스 인수에 주어진 `resolve`와 `reject` 매개변수는 이를 처리하는 데 사용됩니다. 프로미스가 성공하기를 원할 때는 `resolve`를 사용하고, 실패하기를 원할 때는 `reject`를 사용합니다. 이들은 아래에서 보듯 인수를 받는 메서드입니다.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

위 예제는 이 함수들의 인수로 문자열을 사용했지만, 실제로는 무엇이든 될 수 있습니다. 종종 웹사이트나 다른 곳에 표시할 데이터를 담은 객체일 수도 있습니다.

# --instructions--

프로미스가 성공과 실패를 처리하도록 만드세요. 만약 `responseFromServer`가 `true`라면, 프로미스를 성공적으로 완료하기 위해 `resolve` 메서드를 호출하세요. `resolve`에 `We got the data` 값을 가진 문자열을 전달하세요. 만약 `responseFromServer`가 `false`라면, 대신 `reject` 메서드를 사용하고 문자열 `Data not received`을 전달하세요.

# --hints--

`resolve`는 `if` 조건이 `true`일 때 예상 문자열과 함께 호출되어야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

`reject`는 `if` 조건이 `false`일 때 예상 문자열과 함께 호출되어야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
