---
id: 5cddbfd622f1a59093ec611d
title: 모듈 스크립트 만들기
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript는 원래 주로 HTML 웹에서 작은 역할을 하며 시작했습니다. 오늘날에는 매우 커져서 일부 웹사이트는 거의 전적으로 JavaScript로 만들어집니다. JavaScript를 더 모듈화하고 깔끔하며 유지 관리하기 쉽게 만들기 위해 ES6는 JavaScript 파일 간에 코드를 쉽게 공유하는 방법을 도입했습니다. 이는 한 파일의 일부를 내보내 다른 하나 이상의 파일에서 사용하도록 하고, 필요한 부분을 필요한 곳에서 가져오는 것을 포함합니다. 이 기능을 활용하려면 HTML 문서에 `type`가 `module`인 스크립트를 만들어야 합니다. 예시는 다음과 같습니다.

```html
<script type="module" src="filename.js"></script>
```

이 `module` 유형을 사용하는 스크립트는 이제 앞으로 배우게 될 `import`와 `export` 기능을 사용할 수 있습니다.

# --instructions--

HTML 문서에 `module` 유형의 스크립트를 추가하고 소스 파일을 `index.js`로 지정하세요.

# --hints--

`script` 태그를 만들어야 합니다.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

`script` 태그에는 `type` 속성이 있어야 하며, 그 값은 `module`이어야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

`script` 태그에는 `src`인 `index.js`가 있어야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
