---
id: 5cfa550e84205a357704ccb6
title: 객체에서 값을 추출하기 위해 구조 분해 할당 사용하기
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>구조 분해 할당</dfn>은 ES6에서 도입된 특별한 구문으로, 객체에서 직접 가져온 값을 깔끔하게 할당할 수 있습니다.

다음 ES5 코드를 살펴보세요:

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name`는 문자열 `John Doe` 값을 가지며, `age`는 숫자 `34` 값을 갖게 됩니다.

다음은 ES6 구조 분해 구문을 사용한 동등한 할당문입니다:

```js
const { name, age } = user;
```

다시 말해, `name`는 문자열 `John Doe` 값을 가지며, `age`는 숫자 `34` 값을 갖게 됩니다.

여기서 `name`와 `age` 변수는 생성되어 `user` 객체의 각각 해당하는 값으로 할당됩니다. 얼마나 더 깔끔한지 알 수 있습니다.

객체에서 원하는 만큼 많은 값이나 적은 값을 추출할 수 있습니다.

# --instructions--

두 개의 할당문을 동등한 구조 분해 할당으로 바꾸세요. `today`와 `tomorrow` 변수가 `today` 객체에서 `tomorrow`와 `HIGH_TEMPERATURES` 값을 할당받도록 해야 합니다.

# --hints--

ES5 할당 구문을 제거해야 합니다.

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

구조 분해를 사용해 `today` 변수를 생성해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

구조 분해를 사용해 `tomorrow` 변수를 생성해야 합니다.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today`는 `77`와 같아야 하고 `tomorrow`는 `80`와 같아야 합니다.

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
