---
id: 673339d96c56d36e963e888e
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, construirás un catálogo de vivero de plantas mientras practicas los objetos Map y Set de JavaScript.

Para representar cada planta en tu catálogo, usarás objetos, cada uno con tres propiedades:

- `commonName`, el nombre común de la planta
- `scientificName`, el nombre científico (o nombre latino) de la planta
- `cultivar`, la variedad de la planta

Comienza creando un objeto con las tres propiedades anteriores para la `Lavanda española` (nombre científico `Lavandula stoechas`) de un cultivar llamado `Ballerina` y asígnalo a una variable llamada `ballerina`.

# --hints--

Debes tener una variable llamada `ballerina`.

```js
assert.exists(ballerina);
```

Debes asignar un objeto a tu variable `ballerina`.

```js
assert.isObject(ballerina);
```

Tu objeto `ballerina` debe tener una propiedad `commonName` configurada en `"Spanish lavender"`.

```js
assert.strictEqual(ballerina.commonName, "Spanish lavender");
```

Tu objeto `ballerina` debe tener una propiedad `scientificName` configurada en `"Lavandula stoechas"`.

```js
assert.strictEqual(ballerina.scientificName, "Lavandula stoechas");
```

Tu objeto `ballerina` debe tener una propiedad `cultivar` configurada en `"Ballerina"`.

```js
assert.strictEqual(ballerina.cultivar, "Ballerina");
```

Tu objeto `ballerina` debe tener tres claves.

```js
assert.lengthOf(Object.keys(ballerina), 3);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
