---
id: 6733b5e6f210c1bd1afdf8b3
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Los objetos de plantas restantes se han añadido por ti. Ahora es momento de crear el catálogo.

Al igual que los objetos de JavaScript regulares, los objetos `Map` contienen una colección de pares clave-valor. Para crear un objeto `Map`, puedes utilizar el constructor `Map`:

```js
const myMap = new Map();
```

Declara una variable `catalog` y asígnale un objeto `Map`.

# --hints--

Debes tener una variable llamada `catalog`.

```js
assert.exists(catalog);
```

Deberías asignar un objeto `Map` a tu variable `catalog`.

```js
assert.instanceOf(catalog, Map);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--

--fcc-editable-region--
```
