---
id: 6734a13a29acdf47f2b8e161
title: Paso 6
challengeType: 1
dashedName: step-6
---

# --description--

A diferencia de los objetos regulares, donde las claves deben ser siempre cadenas de texto, los objetos de mapa te permiten usar cualquier tipo de dato como claves. En tu catálogo, vas a usar pares de clave-valor objeto-objeto, asociando cada uno de tus objetos de planta con un objeto que contiene el número de macetas disponibles para diferentes tamaños.

Modifica tu llamada a `set` para usar tu objeto `ballerina` como clave. Para el valor, usa un objeto que contenga tres propiedades, `small`, `medium`, y `large`, cada una con el valor de un entero positivo.

# --hints--

Deberías establecer tu objeto `ballerina` como una clave de tu mapa `catalog`.

```js
assert.isTrue(catalog.has(ballerina));
```

El valor de la clave `ballerina` en el mapa `catalog` debería ser un objeto con las propiedades `small`, `medium`, y `large`.

```js
const b = catalog.get(ballerina)
assert.property(b, "small");
assert.property(b, "medium");
assert.property(b, "large");
```

La propiedad `small` del objeto asociado con la clave `ballerina` en el mapa `catalog` debe ser un entero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.small));
assert.isAtLeast(b.small, 0);
```

La propiedad `medium` del objeto asociado con la clave `ballerina` en el mapa `catalog` debe ser un entero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.medium));
assert.isAtLeast(b.medium, 0);
```

La propiedad `large` del objeto asociado con la clave `ballerina` en el mapa `catalog` debe ser un entero positivo.

```js
const b = catalog.get(ballerina)
assert.isTrue(Number.isInteger(b.large));
assert.isAtLeast(b.large, 0);
```

Tu mapa `catalog` debería contener un único par de clave-valor.

```js
assert.lengthOf(catalog, 1);
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();
catalog.set("ballerina", ballerina);
--fcc-editable-region--
console.log(catalog);
console.log(catalog.size);
```
