---
id: 67349ba230cff237c0f2bdf2
title: Paso 4
challengeType: 1
dashedName: step-4
---

# --description--

Los objetos Map recuerdan el orden en que se insertan las claves en el mapa. Para establecer un par clave-valor en tu objeto Map, puedes usar el método `set`. Lo siguiente establece la clave `"name"` con el valor `"Jerry"` en el objeto Map `myMap`.

```js
myMap.set("name", "Jerry");
```

Establece la clave `"ballerina"` con el valor de tu objeto `ballerina` en tu mapa `catalog`. Luego, registra tu catálogo en la consola.

Nota que para ver correctamente tu objeto Map tendrás que abrir la consola del navegador.

# --hints--

Debes configurar la clave `"ballerina"` con el valor de tu objeto `ballerina` en tu mapa `catalog`.

```js
assert.deepEqual(catalog.get("ballerina"), ballerina)
```

Tu mapa `catalog` debe contener un solo par clave-valor.

```js
assert.lengthOf(catalog, 1);
```

Debes registrar tu mapa `catalog` en la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*catalog\s*\)/)
```

# --seed--

## --seed-contents--

```js
const ballerina = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Ballerina"
}

const prettyPolly = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Pretty Polly"
}

const willowVale = {
    commonName: "Spanish lavender",
    scientificName: "Lavandula stoechas",
    cultivar: "Willow Vale"
}

const hidcote = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Hidcote"
}

const imperialGem = {
    commonName: "English lavender",
    scientificName: "Lavandula angustifolia",
    cultivar: "Imperial Gem"
}

const royalCrown = {
    commonName: "French lavender",
    scientificName: "Lavandula dentata",
    cultivar: "Royal Crown"
}
--fcc-editable-region--
const catalog = new Map();

--fcc-editable-region--
```
