---
id: 69899123889dcf5b19be1847
title: 깊이 우선 탐색 알고리즘 구현하기
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

이번 실습에서는 <dfn>깊이 우선 탐색</dfn>(DFS)이라는 그래프 탐색 알고리즘을 구현할 것입니다.

<dfn>너비 우선 탐색</dfn>(BFS)이 이웃 노드를 레벨별로 탐색하는 반면, <dfn>깊이 우선 탐색</dfn>은 한 경로를 따라 가능한 깊게 탐색한 후 되돌아갑니다.

알고리즘이 경로의 끝(방문하지 않은 이웃이 없는 노드)에 도달하면, 아직 탐색하지 않은 간선이 남아 있는 가장 최근 노드로 되돌아가서 탐색을 계속합니다.

이 알고리즘을 반복문으로 구현하려면 **스택**을 사용하세요. JavaScript에서는 표준 `Array`가 `.push()`와 `.pop()` 메서드를 사용해 스택 역할을 하며, **후입선출(LIFO)** 원칙을 따릅니다. 이렇게 하면 스택에 가장 최근에 추가된 이웃 노드가 다음에 탐색됩니다.

이 알고리즘의 간단한 출력은 주어진 노드에서 도달할 수 있는 모든 노드의 목록입니다. 따라서 방문한 노드를 추적하는 것도 필요합니다.

**목표**: 아래 사용자 스토리를 완성하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `dfs`라는 이름의 함수를 정의해야 합니다.
1. `dfs` 함수는 두 개의 인수를 받아야 합니다:

    - `graph`: 인접 행렬(배열의 배열).
    - `root`: 숫자형 노드 레이블(시작 인덱스)로, `0`와 `n - 1` 사이의 노드 숫자이며, `n`는 그래프 내 노드 총 개수입니다.

1. 함수는 시작 노드에서 도달할 수 있는 모든 노드 레이블을 포함하는 배열을 반환해야 합니다.

# --hints--

`dfs`라는 이름의 함수가 두 개의 인수를 받도록 하세요.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)`는 `1`, `2`, `3`, `0`가 포함된 목록을 반환해야 합니다.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)`는 `1`, `2`, `3`, `0`가 포함된 목록을 반환해야 합니다.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)`는 `[3]`를 반환해야 합니다.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)`는 `3`와 `2`가 포함된 목록을 반환해야 합니다.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)`는 `0`와 `1`가 포함된 목록을 반환해야 합니다.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
