---
id: 686d749d6f32f708fdf2b7ad
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In a previous lesson, you learned about `svg` elements and how they're often used as icons in projects. Em bases de código do mundo real, você normalmente confiaria em bibliotecas de ícones, então não precisa criar elementos `svg` do zero. No entanto, neste workshop, você vai construir um ícone de coração para aprender sobre os atributos principais usados dentro de um elemento `svg`.

Comece criando um elemento `svg` na página.

# --hints--

Você deve ter um elemento `svg` na página.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
