---
id: 686daa7ed79ceacd0b264e81
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O ícone de coração está quase pronto.Você só precisa colorir de vermelho.Para fazer isso, defina o atributo `fill` do elemento `svg` como `red`.

Parabéns por concluir este workshop!

# --hints--

Você deve ter um atributo `fill` no seu elemento `svg`.

```js
const svg = document.querySelector('svg');
const fill = svg.getAttribute('fill');
assert.exists(fill);
```

Você deve definir o atributo `fill` para `red`. 

```js
const svg = document.querySelector('svg');
const fill = svg.getAttribute('fill');
assert.strictEqual(fill, "red");
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24" viewBox="0 0 24 24">
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
    --fcc-editable-region--
  </body>
</html>
```

# --solutions--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>

  <body>
    <svg width="24" height="24" viewBox="0 0 24 24" fill="red">
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      />
    </svg>
  </body>
</html>
```
