---
id: 686daa7ed79ceacd0b264e7d
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Você deve aninhar um elemento `path` dentro do seu elemento `svg` para dar forma à imagem.

Crie um elemento `path`.

# --hints--

Você deve ter um elemento `path` aninhado dentro do seu elemento `svg`.

```js
const path = document.querySelector('svg path');
assert.exists(path);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg>
      
    </svg>
    --fcc-editable-region--
  </body>
</html>
```
