---
id: 5f33294a6af5e9188dbdb8f3
title: Passo 1
challengeType: 0
dashedName: step-1
---

# --description--

Neste workshop, você praticará o básico de CSS (Cascading Style Sheets) construindo um menu de café.

Let's start by adding some menu content. Adicione um elemento `main` dentro do elemento `body` existente. Mais pra frente ele conterá informações sobre preços dos cafés e sobremesas oferecidas pela cafeteria.

# --hints--

Seu código deve ter uma tag de abertura `<main>`.

```js
assert.match(code, /<main>/i);
```

Seu código deve ter uma tag de fechamento `</main>`.

```js
assert.match(code, /<\/main>/i);
```

Você não deve alterar seu elemento `body`. Certifique-se de não excluir acidentalmente sua tag de fechamento.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

A tag `main` deve estar dentro da tag `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
