---
id: 5f3313e74582ad9d063e3a38
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Adicione um elemento `head` dentro do elemento `html` para que você possa adicionar um elemento `title`.O texto do elemento `title` deve ser `Cafe Menu`.

# --hints--

Você deve ter uma tag de abertura `<head>`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de fechamento `</head>`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de abertura `<title>`.

```js
assert.match(code, /<title>/i);
```

Você deve ter uma tag de fechamento `</title>`.

```js
assert.match(code, /<\/title>/i);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

Seu elemento `title` deve conter o texto `Cafe Menu`. Você pode precisar verificar sua ortografia.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
