---
id: 5f33071498eb2472b87ddee4
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você aprenderá o básico de CSS (Cascading Style Sheets) construindo um menu de café. 

Comece adicionando seu elemento `<!DOCTYPE html>` e `html` com um atributo `lang` definido como `en`.

# --hints--

Você deve ter o `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter uma abertura `<html lang="en">` com o idioma definido para inglês.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve ter uma tag de fechamento `</html>`.

```js
assert.match(code, /<\/html>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
