---
id: 5f3477aefa51bfc29327200b
title: Passo 10
challengeType: 0
dashedName: step-10
---

# --description--

Você estilizou três elementos escrevendo CSS dentro das tags `style`. Isso funciona, mas como haverá vários outros estilos, é melhor colocar todos os estilos em um arquivo separado e vincular a ele.

Um arquivo `styles.css` separado foi criado para você. Você pode alternar a visualização dos arquivos com as abas no topo do editor.

Comece reescrevendo os estilos que você criou no arquivo `styles.css`. Certifique-se de excluir a tag de abertura e fechamento do `style`.

# --hints--

Seu arquivo `styles.css` deve ter o seletor de tipo `h1, h2, p`.

```js
assert.match(editableContents.replace(/[\s\n]*/g, ""), /(h1|h2|p),(h1|h2|p),(h1|h2|p){/);
```
 
O seletor deve definir a propriedade `text-align` para `center`.

```js
assert.match(editableContents, /text-align:\s*center;?/);
```

Seu código não deve conter as tags `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
