---
id: 5f344fc1520b6719f2e35605
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

No menu, haverá duas seções, uma para cafés e outra para sobremesas. Adicione um elemento `section` dentro do elemento `main` para que você tenha um lugar onde colocar todos os cafés disponíveis.

# --hints--

Você deve ter uma tag de abertura `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

Você deve ter uma tag de fechamento `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

Você não deve alterar seu elemento `main` existente. Certifique-se de que você não tenha excluído a tag de fechamento.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

O elemento `section` deve estar dentro do elemento `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
