---
id: 5f3477cbcb6ba47918c1da92
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Para que a estilização da página pareça semelhante no celular, em um desktop ou em um laptop, você precisa adicionar um elemento `meta` com um atributo especial `content`.

Você aprendeu sobre o elemento viewport `meta` nas lições anteriores.

Aqui está um exemplo:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

O código deve ter dois elementos `meta`.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

O elemento `meta` deve ter um atributo `name` com o valor `viewport`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

O elemento `meta` deve ter um atributo `content` com o valor `width=device-width, initial-scale=1.0`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
