---
id: 66bb6a9c2dd58b73cd759034
title: Eine Notenbuch-App erstellen
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion namens `getAverage` haben, die ein Array von Testergebnissen als Parameter entgegennimmt und die durchschnittliche Punktzahl zurückgibt.
1. Sie sollten eine Funktion namens `getGrade` haben, die eine Schülerpunktzahl als Parameter entgegennimmt und einen String zurückgibt, der eine Note basierend auf der Punktzahl darstellt. Hier sind die Punktzahlen und ihre entsprechenden Noten:

    | Punktzahlbereich | Note   |
    | ----------- | ------- |
    | `100`       | `"A+"`  |
    | `90 - 99`   | `"A"`   |
    | `80 - 89`   | `"B"`   |
    | `70 - 79`   | `"C"`   |
    | `60 - 69`   | `"D"`   |
    | `0 - 59`    | `"F"`   |

1. Sie sollten eine Funktion namens `hasPassingGrade` haben, die eine Punktzahl als Parameter entgegennimmt und entweder `true` oder `false` zurückgibt, je nachdem, ob die Punktzahl einer bestandenen Note entspricht.
1. Die Funktion `hasPassingGrade` sollte die Funktion `getGrade` verwenden, um die Note zu erhalten, und damit bestimmen, ob die Note bestanden ist. Eine bestandene Note ist alles, was nicht `"F"` ist.
1. Sie sollten eine Funktion namens `studentMsg` haben, die ein Array von Punktzahlen und eine Schülerpunktzahl als Parameter entgegennimmt. Die Funktion sollte einen String im folgenden Format zurückgeben:

    - `"Class average: average-goes-here. Your grade: grade-goes-here. You passed the course."`, wenn der Schüler den Kurs bestanden hat.
    - `"Class average: average-goes-here. Your grade: grade-goes-here. You failed the course."`, wenn der Schüler den Kurs nicht bestanden hat.

    Ersetzen Sie `average-goes-here` durch den Durchschnitt der Gesamtpunktzahlen und `grade-goes-here` durch die Note des Schülers.
    Verwenden Sie `getAverage`, um die durchschnittliche Punktzahl zu erhalten, und `getGrade`, um die Note des Schülers zu erhalten.

# --hints--

Sie sollten eine Funktion namens `getAverage` haben.

```js
assert.isFunction(getAverage);
```

Ihre Funktion `getAverage` sollte eine Zahl zurückgeben.

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` sollte `71.7` zurückgeben.

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` sollte `85.4` zurückgeben.

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` sollte `92.4` zurückgeben.

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` sollte `47.25` zurückgeben.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` sollte `47.125` zurückgeben.

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

Ihre Funktion `getAverage` sollte die durchschnittliche Punktzahl der Testergebnisse zurückgeben.

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

Sie sollten eine Funktion namens `getGrade` haben.

```js
assert.isFunction(getGrade);
```

Ihre Funktion `getGrade` sollte `"A+"` zurückgeben, wenn die Punktzahl `100` ist.

```js
assert.strictEqual(getGrade(100), "A+");
```

Ihre Funktion `getGrade` sollte `"A"` zurückgeben, wenn die Punktzahl zwischen `90` und `99` liegt.

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

Ihre Funktion `getGrade` sollte `"B"` zurückgeben, wenn die Punktzahl zwischen `80` und `89` liegt.

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

Ihre Funktion `getGrade` sollte `"C"` zurückgeben, wenn die Punktzahl zwischen `70` und `79` liegt.

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

Ihre Funktion `getGrade` sollte `"D"` zurückgeben, wenn die Punktzahl zwischen `60` und `69` liegt.

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

Ihre Funktion `getGrade` sollte `"F"` zurückgeben, wenn die Punktzahl zwischen `0` und `59` liegt.

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


Sie sollten eine Funktion namens `hasPassingGrade` haben.

```js
assert.isFunction(hasPassingGrade);
```

Ihre Funktion `hasPassingGrade` sollte einen Boolean-Wert zurückgeben.

```js
assert.isBoolean(hasPassingGrade(100));
```

Ihre Funktion `hasPassingGrade` sollte `true` zurückgeben, wenn die Note ein `"A"` ist.

```js
assert.isTrue(hasPassingGrade(100));
```

Ihre Funktion `hasPassingGrade` sollte `false` zurückgeben, wenn die Note ein `"F"` ist.

```js
assert.isFalse(hasPassingGrade(53));
```

Ihre Funktion `hasPassingGrade` sollte `true` für alle bestandenen Noten zurückgeben.

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

Sie sollten eine Funktion namens `studentMsg` haben.

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` sollte die folgende Nachricht zurückgeben: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` sollte die folgende Nachricht zurückgeben: `"Class average: 50.8. Your grade: A+. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` sollte die folgende Nachricht zurückgeben: `"Class average: 48.25. Your grade: B. You passed the course."`.

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` sollte die folgende Nachricht zurückgeben: `"Class average: 45.625. Your grade: C. You passed the course."`.

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

Ihre Funktion `studentMsg` sollte die korrekte Nachricht basierend auf der Punktzahl des Schülers und dem Klassendurchschnitt zurückgeben.

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
