---
id: 6976db7ecfa770bf21307b20
title: 建置播放清單混音引擎
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

在這個實驗中，你將建置一個程式，從多個聽眾提交的播放清單中創建一個混音播放清單。

每位聽眾會提供他們想聽的歌曲列表。有些歌曲可能出現多次，有些歌手也可能出現過多次。你的工作是逐步處理這些播放清單：將它們合併成一個列表，為每首歌評分，移除重複歌曲，限制同一歌手出現的次數，然後建立最終的播放順序。

**目標：** 完成以下使用者故事並通過所有測試以完成此實驗。

**使用者故事：**

1. 你應該創建一個名為 `flattenPlaylists` 的函式，接受一個播放清單陣列，其中每個播放清單是包含以下屬性的物件陣列：`trackId`、`artist`、`title`、`votes`、`bpm`。如果輸入不是陣列，`flattenPlaylists` 應該傳回空陣列。已為你提供一個範例播放清單，你可以用它來測試你的函式。

2. `flattenPlaylists` 應該傳回一個扁平化的曲目物件陣列，每個物件包含所有原始曲目屬性，外加一個 `source` 屬性，該屬性是包含播放清單索引和曲目索引的陣列，用來表示曲目的來源。

3. 你應該創建一個名為 `scoreTracks` 的函式，接受由 `flattenPlaylists` 傳回的曲目物件陣列（每個物件包含 `trackId`、`artist`、`title`、`votes`、`bpm` 和 `source` 屬性），並傳回一個新的曲目物件陣列，每個物件新增一個 `score` 屬性，使用公式 `votes * 10 - Math.abs(bpm - 120)` 計算。

4. 你應該創建一個名為 `dedupeTracks` 的函式，接受由 `scoreTracks` 傳回的曲目物件陣列，並傳回一個新的陣列，移除重複的 `trackId` 項目，只保留第一次出現的。

5. 你應該創建一個名為 `enforceArtistQuota` 的函式，接受由 `dedupeTracks` 傳回的曲目物件陣列和一個數字，代表每位歌手允許的最大出現次數。該函式應傳回一個新的陣列，確保沒有歌手出現超過該數字，並保留最早出現的曲目。

6. 你應該創建一個名為 `buildSchedule` 的函式，接受由 `enforceArtistQuota` 傳回的曲目物件陣列，並傳回一個新的 `{ slot, trackId }` 物件陣列，其中 `slot` 是以 1 為起點的索引，代表每首曲目在播放順序中的位置。

7. 你應該創建一個名為 `remixPlaylist` 的函式，接受播放清單陣列和每位歌手允許的最大出現次數。該函式應透過依序呼叫 `{ slot, trackId }`、`flattenPlaylists`、`scoreTracks`、`dedupeTracks` 和 `enforceArtistQuota`，傳回最終的播放排程，為 `buildSchedule` 物件陣列。

# --hints--

你應該有一個名為 `flattenPlaylists` 的函式。

```js
assert.isFunction(flattenPlaylists);
```

當輸入不是陣列時，`flattenPlaylists` 應該傳回空陣列。

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

`flattenPlaylists` 傳回的每個曲目應包含一個 `source` 欄位，該欄位是包含播放清單索引和曲目索引的陣列。

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

你應該有一個名為 `scoreTracks` 的函式。

```js
assert.isFunction(scoreTracks)
```

`scoreTracks` 傳回的每個曲目應包含一個數字型別的 `score` 欄位。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

你應該使用目標 BPM 為 `score`，並用公式 `120` 計算 `votes * 10 - Math.abs(bpm - 120)`。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

你應該有一個名為 `dedupeTracks` 的函式。

```js
assert.isFunction(dedupeTracks)
```

當存在重複的 `trackId` 值時，`dedupeTracks` 應只保留第一次出現的曲目。

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

你應該有一個名為 `enforceArtistQuota` 的函式。

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` 應確保沒有歌手出現超過 `maxPerArtist` 次，透過移除多餘曲目並保留最早出現的。

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

你應該有一個名為 `buildSchedule` 的函式。

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` 應傳回形狀為 `{ slot, trackId }` 的物件陣列，其中 `slot` 從 `1` 開始。

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

你應該有一個名為 `remixPlaylist` 的函式。

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` 應依序呼叫輔助函式以產生最終排程。

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
