---
id: 612ea4c4993aba52ab4aa32e
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Jetzt ist es Zeit, die Pseudo-Selektoren zu verwenden, die Sie zuvor vorbereitet haben. Um die schwarzen Tasten zu erstellen, fügen Sie einen neuen `.key.black--key::after`-Selektor hinzu. Dieser richtet sich an die Elemente mit der Klasse `key black--key` und wählt das Pseudoelement nach diesen Elementen im HTML aus.

Setzen Sie im neuen Selektor die `background-color`-Eigenschaft auf `#1d1e22`. Setzen Sie außerdem die `content`-Eigenschaft auf `""`. Dadurch werden die Pseudoelemente leer.

Die `content`-Eigenschaft wird verwendet, um den Inhalt des Elements festzulegen oder zu überschreiben. Standardmäßig sind die Pseudoelemente, die durch die `::before`- und `::after`-Pseudo-Selektoren erstellt werden, leer, und die Elemente werden nicht auf der Seite gerendert. Wenn Sie die `content`-Eigenschaft auf einen leeren String `""` setzen, wird sichergestellt, dass das Element auf der Seite gerendert wird, obwohl es leer bleibt.

Wenn Sie experimentieren möchten, versuchen Sie, die `background-color`-Eigenschaft zu entfernen und verschiedene Werte für die `content`-Eigenschaft zu setzen, wie zum Beispiel `"♥"`. Denken Sie daran, diese Änderungen rückgängig zu machen, wenn Sie fertig sind, damit die Tests bestehen.

# --hints--

Sie sollten einen `.key.black--key::after`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after'));
```

Ihr `.key.black--key::after`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `#1d1e22` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.backgroundColor === 'rgb(29, 30, 34)');
```

Ihr `.key.black--key::after`-Selektor sollte eine `content`-Eigenschaft mit dem Wert `""` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key.black--key::after')?.content === '""');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

.key {
  background-color: #ffffff;
  position: relative;
  width: 41px;
  height: 175px;
  margin: 2px;
  float: left;
}

--fcc-editable-region--

--fcc-editable-region--
```
