---
id: 612e9d142affc44a453655db
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Es ist Zeit, die Tasten selbst zu gestalten. Erstellen Sie einen `class`-Selektor für die `.key`-Elemente. Setzen Sie die `background-color`-Eigenschaft auf den Wert `#ffffff`, die `position`-Eigenschaft auf `relative`, die `width`-Eigenschaft auf `41px` und die `height`-Eigenschaft auf `175px`.

# --hints--

Sie sollten einen `.key`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key'));
```

Ihr `.key`-Selektor sollte eine `background-color`-Eigenschaft mit dem Wert `#ffffff` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.backgroundColor === 'rgb(255, 255, 255)');
```

Ihr `.key`-Selektor sollte die `position`-Eigenschaft auf `relative` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.position === 'relative');
```

Ihr `.key`-Selektor sollte eine `width`-Eigenschaft mit dem Wert `41px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.width === '41px');
```

Ihr `.key`-Selektor sollte eine `height`-Eigenschaft mit dem Wert `175px` haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('.key')?.height === '175px');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
html {
  box-sizing: border-box;
}

*, *::before, *::after {
  box-sizing: inherit;
}

#piano {
  background-color: #00471b;
  width: 992px;
  height: 290px;
  margin: 80px auto;
  padding: 90px 20px 0 20px;
}

.keys {
  background-color: #040404;
  width: 949px;
  height: 180px;
  padding-left: 2px;
}

--fcc-editable-region--

--fcc-editable-region--
```
