---
id: 612e77aba7ca691f598feb02
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie zwei `meta`-Tags hinzu, einen, um Ihre Seite für mobile Geräte zu optimieren, und einen, um ein akzeptiertes `charset` für die Seite anzugeben.

# --hints--

Sie sollten zwei `meta`-Elemente zu Ihrer Seite hinzufügen.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Sie sollten zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Ein `meta`-Element sollte ein `name` mit dem Wert `viewport` und ein `content` mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut mit dem Wert `UTF-8` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Piano</title>
    --fcc-editable-region--

    --fcc-editable-region--
  </head>
  <body></body>
</html>
```

```css

```
