---
id: 612e89d254fe5d3df7d6693d
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Nachdem Sie das `html`-Boxmodell zurückgesetzt haben, müssen Sie dies auch an die darin enthaltenen Elemente weitergeben. Dazu können Sie die `box-sizing`-Eigenschaft auf `inherit` setzen, was den Ziel-Elementen mitteilt, denselben Wert wie das Elternelement zu verwenden.

Sie müssen auch die Pseudoelemente ansprechen, das sind spezielle Schlüsselwörter, die einem Selektor folgen. Die beiden Pseudoelemente, die Sie verwenden werden, sind die `::before`- und `::after`-Pseudoelemente.

Der `::before`-Selektor erzeugt ein Pseudoelement, das das erste Kind des ausgewählten Elements ist, während der `::after`-Selektor ein Pseudoelement erzeugt, das das letzte Kind des ausgewählten Elements ist. Diese Pseudoelemente werden oft verwendet, um kosmetische Inhalte zu erstellen, was Sie später in diesem Projekt sehen werden.

Erstellen Sie jetzt einen CSS-Selektor, der alle Elemente mit `*` anspricht, und schließen Sie die Pseudoelemente mit `::before` und `::after` ein. Setzen Sie die `box-sizing`-Eigenschaft auf `inherit`.

# --hints--

Sie sollten einen `*, ::before, ::after`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after'));
```

Ihr `*, ::before, ::after`-Selektor sollte die `box-sizing`-Eigenschaft auf `inherit` gesetzt haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('*, ::before, ::after')?.boxSizing === 'inherit');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
html {
  box-sizing: border-box;
}
--fcc-editable-region--
```
