---
id: 612e6afc009b450a437940a1
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Beginnen Sie mit der grundlegenden HTML-Struktur. Fügen Sie eine `DOCTYPE` Deklaration sowie `html`, `head`, `body` und `title` Elemente hinzu.

Legen Sie die Sprache dieser Seite auf Englisch fest. Setzen Sie das `title` auf `Piano`.

# --hints--

Ihr Code sollte die `DOCTYPE` Referenz enthalten.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Sie sollten nach der `DOCTYPE` Referenz ein Leerzeichen einfügen.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Sie sollten den Dokumenttyp als `html` definieren.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Sie sollten die `DOCTYPE` Deklaration mit einem `>` nach dem Typ schließen.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Ihr `html` Element sollte ein öffnendes Tag mit einem `lang` Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html` Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/));
```

Ihre `DOCTYPE` Deklaration sollte am Anfang Ihres HTML stehen.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Sie sollten ein öffnendes `head` Tag haben.

```js
assert(code.match(/<head\s*>/i));
```

Sie sollten ein schließendes `head` Tag haben.

```js
assert(code.match(/<\/head\s*>/i));
```

Sie sollten ein öffnendes `body` Tag haben.

```js
assert(code.match(/<body\s*>/i));
```

Sie sollten ein schließendes `body` Tag haben.

```js
assert(code.match(/<\/body\s*>/i));
```

Die `head` und `body` Elemente sollten Geschwister sein.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Das `head` Element sollte innerhalb des `html` Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Das `body` Element sollte innerhalb des `html` Elements sein.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Ihr Code sollte ein `title` Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr Projekt sollte einen Titel von `Piano` haben.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Denken Sie daran, dass Groß- und Kleinschreibung sowie Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
