---
id: 612e813b3ba67633222cbe54
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Denken Sie daran, dass ein `class`-Attribut mehrere Werte haben kann. Um Ihre weißen Tasten von Ihren schwarzen Tasten zu trennen, fügen Sie einen zweiten `class`-Wert von `black--key` hinzu. Fügen Sie dies zu Ihrem zweiten, dritten, fünften, sechsten und siebten `.key`-Element hinzu.

# --hints--

Ihr zweites `.key`-Element sollte ebenfalls ein `class` von `black--key` haben.

```js
const key = document.querySelectorAll('.key')?.[1];
assert(key?.className?.includes('black--key'));
```

Ihr drittes `.key` sollte `black--key` im `class` haben.

```js
const third = document.querySelectorAll('.key')?.[2];
assert(third?.classList?.contains('black--key'));
```

Ihr fünftes `.key` sollte `black--key` im `class` haben.

```js
const fifth = document.querySelectorAll('.key')?.[4];
assert(fifth?.classList?.contains('black--key'));
```

Ihr sechstes `.key` sollte `black--key` im `class` haben.

```js
const sixth = document.querySelectorAll('.key')?.[5];
assert(sixth?.classList?.contains('black--key'));
```

Ihr siebtes `.key` sollte `black--key` im `class` haben.

```js
const seventh = document.querySelectorAll('.key')?.[6];
assert(seventh?.classList?.contains('black--key'));
```

Sie sollten fünf `.black--key`-Elemente haben.

```js
const blackKeys = document.querySelectorAll('.black--key');
assert(blackKeys?.length === 5);
```

Sie sollten sieben `.key`-Elemente haben.

```js
const keys = document.querySelectorAll('.key');
assert(keys?.length === 7);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
    --fcc-editable-region--
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
        <div class="key"></div>
      </div>
    </div>
    --fcc-editable-region--
  </body>
</html>
```

```css

```
