---
id: 685976a92ffab8b2dc744b72
title: Passo 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro do seu elemento `.tabs`, adicione um elemento `div` com um atributo `role` definido como `"tablist"` e um atributo `aria-labelledby` definido como `"tabs-title"`.  

O atributo `aria-labelledby` é necessário aqui para associar o tablist a um rótulo visível, ajudando tecnologias assistivas como leitores de tela a anunciar o propósito do tablist para os usuários.

# --hints--

Você deve ter um elemento `div` dentro do seu elemento `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Seu elemento `div` deve ter um atributo `role` definido como `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Seu elemento `div` deve ter um atributo `aria-labelledby` definido como `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
