---
id: 68599af28ead38c22590a89d
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Deve haver uma maneira de conectar os botões ao conteúdo que eles representam.Isso é feito usando o atributo `aria-controls` em cada botão.Mais tarde, você adicionará o conteúdo e os atributos `id` correspondentes aos elementos de conteúdo.

Comece adicionando o atributo `aria-controls` a cada um dos elementos `button`.O valor para o atributo `aria-controls` do primeiro botão deve ser definido como `"panel-earth"`.O atributo `aria-controls` do segundo botão deve ser definido como `"panel-saturn"`.O atributo `aria-controls` do terceiro botão deve ser definido como `"panel-mars"`.

# --hints--

Cada um dos seus elementos `button` deve ter um atributo `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

O primeiro elemento `button` deve ter um atributo `aria-controls` definido como `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

O segundo elemento `button` deve ter um atributo `aria-controls` definido como `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

O terceiro elemento `button` deve ter um atributo `aria-controls` definido como `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
