---
id: 68599f9df77791c4a311cc5f
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

O último item a ser adicionado a cada um dos elementos `button` são os atributos `id`.Estes serão usados para vincular os elementos `button` aos seus respectivos painéis de conteúdo.

O primeiro elemento `button` deve ter um `id` de `"tab-earth"`, o segundo deve ter um `id` de `"tab-saturn"` e o terceiro deve ter um `id` de `"tab-mars"`.

# --hints--

Seu primeiro elemento `button` deve ter um atributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("id")
);
```

Seu primeiro elemento `button` deve ter um atributo `id` definido como `"tab-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("id"), "tab-earth");
```

Seu segundo elemento `button` deve ter um atributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("id")
);
```

Seu segundo elemento `button` deve ter um atributo `id` definido como `"tab-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("id"), "tab-saturn");
```

Seu terceiro elemento `button` deve ter um atributo `id`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("id")
);
```

Seu terceiro elemento `button` deve ter um atributo `id` definido como `"tab-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("id"), "tab-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
