---
id: 6859ae3ef795c4c9d4badf9c
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora é hora de começar a adicionar a funcionalidade à tablist. O primeiro passo é acessar todos os elementos de aba e painel na `tablist`. 

Crie uma variável chamada `tabs` e atribua a ela o resultado da consulta por todos os elementos com o atributo `role` definido como `tab`.

Então, crie uma variável chamada `panels` e atribua a ela o resultado da consulta de todos os elementos com o atributo `role` definido como `tabpanel`.

Lembre-se de que você pode usar o método `querySelectorAll` para selecionar múltiplos elementos de uma vez.

# --hints--

Você deve ter uma variável chamada `tabs`.

```js
assert.isDefined(tabs);
```

Você deve atribuir `document.querySelectorAll('[role="tab"]')` à variável `tabs`.

```js
assert.deepEqual(
  Array.from(tabs),
  Array.from(document.querySelectorAll('[role="tab"]'))
);
```

Você deve ter uma variável chamada `panels`.

```js
assert.isDefined(panels);
```

Você deve atribuir `document.querySelectorAll('[role="tabpanel"]')` à variável `panels`.

```js
assert.deepEqual(
  Array.from(panels),
  Array.from(document.querySelectorAll('[role="tabpanel"]'))
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      <div id="panel-saturn" role="tabpanel" aria-labelledby="tab-saturn" hidden>
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div id="panel-mars" role="tabpanel" aria-labelledby="tab-mars" hidden>
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
