---
id: 685db0bd1280154ae85bcf9f
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Neste workshop, você irá construir uma tablist de planetas acessível usando HTML, CSS e JavaScript.O boilerplate HTML e o CSS foram fornecidos para você.

Nesses primeiros passos, você irá construir o restante do HTML para o tablist.

Comece criando um elemento `div` com uma classe definida como `"tabs"`.Dentro desse elemento `div`, adicione um elemento `h2` com um atributo `id` definido como `"tabs-title"` e o conteúdo de texto `Planets`.

# --hints--

Você deve ter um elemento `div`.

```js
assert.exists(document.querySelector("div"));
```

Seu elemento `div` deve ter uma classe definida como `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Seu elemento `div` deve ter um elemento `h2` dentro dele.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Seu elemento `h2` deve ter um atributo `id` definido como `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Seu elemento `h2` deve ter o conteúdo de texto `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
