---
id: 687ad7be1ff45032ccf1d92f
title: 테마 스위처 만들기
challengeType: 25
dashedName: lab-theme-switcher
demoType: onClick
---

# --description--

이 실습에서는 서로 다른 테마를 전환하는 앱을 만듭니다. 테마를 바꾸면 배경색이 바뀌고 다른 메시지가 표시됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `button` 요소에 `Switch Theme` 텍스트가 있어야 합니다.
2. `button` 요소에는 다음 속성이 있어야 합니다:
   - `id` 속성이 `"theme-switcher-button"`로 설정되어야 합니다.
   - `aria-haspopup` 속성이 `"true"`로 설정되어야 합니다.
   - `aria-expanded` 속성이 `"false"`로 설정되어야 합니다.
   - `aria-controls` 속성이 `"theme-dropdown"`로 설정되어야 합니다.
3. `ul` 요소에는 다음 속성이 있어야 합니다:
   - `id` 속성이 `"theme-dropdown"`로 설정되어야 합니다.
   - `role` 속성이 `"menu"`로 설정되어야 합니다.
   - `aria-labelledby` 속성이 `"theme-switcher-button"`로 설정되어야 합니다.
   - `hidden` 속성이 있어야 합니다.
4. `ul` 요소에는 최소 두 개의 `li` 요소가 있어야 하며, 각각 `role` 속성이 `"menuitem"`로 설정되어 있고, 서로 다른 테마를 나타내는 텍스트가 있어야 합니다.
5. 각 `li` 요소에는 `id` 속성이 있어야 하며, 이 속성은 `theme-`로 시작하고 `li` 요소 텍스트에 설정한 테마로 끝나야 합니다. 예를 들어, 테마 중 하나가 `Light`이라면 `id`은 `theme-light`이어야 합니다.
6. `aria-live` 속성이 `"polite"`로 설정된 요소가 있어야 합니다.
7. `themes` 배열에는 최소 두 개의 객체가 있어야 하며, 각각 `name`과 `message` 속성을 포함해야 합니다. `name`는 서로 다른 테마를 나타내고 `message`은 테마가 전환될 때 표시됩니다. `name`과 `message` 값은 자유롭게 정할 수 있지만 `name` 값은 `li` 요소의 텍스트 내용과 일치해야 합니다.
8. 사용자가 `#theme-switcher-button`을 클릭하면 `#theme-dropdown` 요소에 사용 가능한 테마가 표시되어야 합니다.
9. `#theme-dropdown` 요소가 표시되면 `aria-expanded` 속성이 `"true"`로 설정되어 `button` 요소에 적용되어야 합니다.
10. 사용자가 `#theme-switcher-button`을 클릭하고 `#theme-dropdown` 요소가 표시된 상태라면 메뉴가 닫히고 `hidden` 속성이 추가되며 `aria-expanded` 속성이 `"false"`로 설정되어야 합니다.
11. 사용자가 `#theme-switcher-button`를 클릭해 테마를 선택하면, 해당 `theme-<name>` 클래스가 `body` 요소에 추가되고 `themes` 배열에 설정한 관련 테마 메시지가 `aria-live="polite"` 요소에 표시되어야 합니다.
12. 사용자가 `#theme-switcher-button`를 클릭해 현재와 다른 테마를 선택하면, `body`의 테마 클래스와 상태 메시지가 새 선택에 맞게 업데이트되어야 합니다.

# --before-all--

```js
const resetThemeSwitch = () => {
  const themeDropdownBtn = document.getElementById("theme-switcher-button");
  const themeDropdownEl = document.getElementById("theme-dropdown");

  // Reset attributes due to earlier button clicks interfering with these tests
  themeDropdownEl.setAttribute("hidden", "");
  themeDropdownBtn.setAttribute("aria-expanded", "false");
}
```

# --hints--

`button` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("button"));
```

`button` 요소에 `Switch Theme` 텍스트가 있어야 합니다.

```js
const btn = document.querySelector("button");
assert.equal(btn?.textContent.trim(), "Switch Theme");
```

`button` 요소에 `id` 속성이 `"theme-switcher-button"`로 설정되어 있어야 합니다.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("id"), "theme-switcher-button");
```

`button` 요소에 `aria-haspopup` 속성이 `"true"`로 설정되어 있어야 합니다.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-haspopup"), "true");
```

`button` 요소에 `aria-expanded` 속성이 `"false"`로 설정되어 있어야 합니다.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-expanded"), "false");
```

`button` 요소에 `aria-controls` 속성이 `"theme-dropdown"`로 설정되어 있어야 합니다.

```js
const btn = document.querySelector("button");
assert.equal(btn?.getAttribute("aria-controls"), "theme-dropdown");
```

`ul` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("ul"));
```

`ul` 요소에 `id` 속성이 `"theme-dropdown"`로 설정되어 있어야 합니다.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("id"), "theme-dropdown");
```

`ul` 요소에 `role` 속성이 `"menu"`로 설정되어 있어야 합니다.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("role"), "menu");
```

`ul` 요소에 `aria-labelledby` 속성이 `"theme-switcher-button"`로 설정되어 있어야 합니다.

```js
const ul = document.querySelector("ul");
assert.equal(ul?.getAttribute("aria-labelledby"), "theme-switcher-button");
```

`ul` 요소에 `hidden` 속성이 있어야 합니다.

```js
const ul = document.querySelector("ul");
assert.isTrue(ul?.hasAttribute("hidden"));
```

`ul` 요소에 최소 두 개의 `li` 요소가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul li");
assert.isAtLeast(listItems?.length, 2);
```

각 `li` 요소에 `role` 속성이 `"menuitem"`로 설정되어 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  assert.equal(li?.getAttribute("role"), "menuitem");
})
```

각 `li` 요소에는 테마를 나타내는 텍스트가 있어야 합니다.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
listItems.forEach(li => {
  const text = li?.textContent.trim();
  assert.isTrue(text.length > 0);
})
```

각 `li` 요소에는 `id` 속성이 있어야 하며, 이 속성은 `theme-`로 시작하고 해당 요소의 테마로 끝나야 합니다. `id` 값은 모두 소문자여야 합니다.

```js
const listItems = document.querySelectorAll("ul li");
assert.isNotEmpty(listItems);
const listTitles = [...listItems].map(li => li?.textContent.trim().toLowerCase());

listItems.forEach((li, index) =>{
  assert.equal(li.getAttribute("id"), `theme-${listTitles[index]}`);
})
```

`aria-live` 속성이 `"polite"`로 설정된 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('[aria-live="polite"]'));
```

`themes` 배열이 있어야 합니다.

```js
assert.isArray(themes);
```

`themes` 배열에는 최소 두 개의 객체가 있어야 합니다.

```js
assert.isAtLeast(themes?.length, 2);
themes.forEach(theme => {
  assert.isObject(theme);
});
```

`themes` 배열에는 `name`와 `message` 속성을 가진 최소 두 개의 객체가 있어야 합니다.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.property(theme, "name");
  assert.property(theme, "message");
})
```

`name`와 `message` 속성 값은 비어 있지 않은 문자열이어야 합니다.

```js
assert.isNotEmpty(themes);
themes?.forEach(theme => {
  assert.isObject(theme);
  assert.isString(theme["name"]);
  assert.isNotEmpty(theme["name"]);
  assert.isString(theme["message"]);
  assert.isNotEmpty(theme["message"]);
})
```

각 `name` 속성 값은 `li` 요소 중 하나의 텍스트 내용과 일치해야 합니다. 이 값들은 모두 소문자여야 합니다.

```js
const listItems = [...document.querySelectorAll("ul li")]?.map(li => li.textContent.toLowerCase().trim());
const namesArr = themes?.map(({name}) => name);

assert.equal(listItems.length, namesArr.length);
assert.isTrue(namesArr?.every(name => listItems.includes(name)));
```

사용자가 `#theme-switcher-button`를 클릭해 테마 옵션을 표시하면 `hidden` 속성이 `#theme-dropdown` 요소에서 제거되어야 합니다.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
```

`#theme-dropdown` 요소가 표시되면 `aria-expanded` 속성이 `"true"`로 설정되어 `button` 요소에 적용되어야 합니다.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
```

사용자가 `#theme-switcher-button`를 클릭하고 `#theme-dropdown` 요소가 표시된 상태라면 드롭다운 메뉴가 `hidden`로 설정되어야 합니다.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const themeDropdownEl = document.getElementById("theme-dropdown");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isFalse(themeDropdownEl?.hasAttribute("hidden"));
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.isTrue(themeDropdownEl?.hasAttribute("hidden"));
```

사용자가 `#theme-switcher-button`를 클릭하고 `#theme-dropdown` 요소가 표시된 상태라면 `aria-expanded` 속성이 `"false"`로 설정되어야 합니다.

```js
resetThemeSwitch()
const themeDropdownBtn = document.getElementById("theme-switcher-button");

themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "true");
themeDropdownBtn?.dispatchEvent(new Event("click"));
assert.equal(themeDropdownBtn?.getAttribute("aria-expanded"), "false");
```

사용자가 `#theme-switcher-button`를 클릭해 테마를 선택하면 해당 `theme-<name>` 클래스가 `body` 요소에 추가되어야 합니다.

```js
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn.dispatchEvent(new Event("click"));

const clickedItem = listItems[1];
clickedItem?.dispatchEvent(new Event("click", {bubbles: true}));

const expectedClass = clickedItem?.id;
const actualClass = document.body.className;

assert.equal(actualClass, expectedClass);
```

사용자가 `#theme-switcher-button`를 클릭해 테마를 선택하면 `message` 배열에서 선택한 `theme`와 관련된 `themes`이 `aria-live="polite"` 요소에 표시되어야 합니다.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("ul li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", {bubbles: true}));

const selectedThemeName = listItems[1]?.textContent.toLowerCase().trim();
const selectedTheme = themes?.find(theme => theme.name === selectedThemeName);

assert.equal(liveRegion?.textContent, selectedTheme?.message);
```

사용자가 `#theme-switcher-button`를 클릭해 서로 다른 두 테마를 선택하면, `body` 클래스와 상태 메시지가 각각 선택한 테마에 맞게 업데이트되어야 합니다.

```js
const liveRegion = document.querySelector('[aria-live="polite"]');
const themeDropdownBtn = document.getElementById("theme-switcher-button");
const listItems = document.querySelectorAll("#theme-dropdown li");

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[0]?.dispatchEvent(new Event("click", { bubbles: true }));
const firstClass = document.body.className;
const firstMessage = liveRegion?.textContent;

themeDropdownBtn?.dispatchEvent(new Event("click"));
listItems[1]?.dispatchEvent(new Event("click", { bubbles: true }));
const secondClass = document.body.className;
const secondMessage = liveRegion?.textContent;

assert.equal(firstClass, listItems[0]?.id);
assert.equal(secondClass, listItems[1]?.id);
assert.notEqual(firstClass, secondClass);
assert.notEqual(firstMessage, secondMessage);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body>

  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

#status {
  text-align: center;
  min-height: 20px;
}
```

```js
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Theme Switcher</title>
  <link rel="stylesheet" href="./styles.css">
</head>
<body class="theme-light">  
  <button
    class="dropdown-btn"
    id="theme-switcher-button"
    aria-haspopup="true"
    aria-expanded="false"
    aria-controls="theme-dropdown"
  >
    Switch Theme
  </button>

  <ul
    class="dropdown-content"
    id="theme-dropdown"
    role="menu"
    aria-labelledby="theme-switcher-button"
    hidden
  >
    <li class="dropdown-item" id="theme-light" role="menuitem">Light</li>
    <li class="dropdown-item" id="theme-dark" role="menuitem">Dark</li>
    <li class="dropdown-item" id="theme-ocean" role="menuitem">Ocean</li>
    <li class="dropdown-item" id="theme-nord" role="menuitem">Nord</li>
  </ul>
  
  <div id="status" aria-live="polite"></div>
  
  <script src="./script.js"></script>
</body>
</html>
```

```css
body {
  margin: 0;
  font-family: sans-serif;
  transition: background 0.3s, color 0.3s;
}

ul {
  margin: 0;
  padding: 0;
}

li {
  list-style-type: none;
}

.dropdown-btn {
  position: relative;
  padding: 10px 20px;
  border: none;
  cursor: pointer;
}

.dropdown-content {
  position: absolute;
  background-color: #ffffff;
  min-width: 160px;
  box-shadow: 0px 4px 8px rgba(0,0,0,0.2);
  z-index: 1;
}

.dropdown-item {
  width: 100%;
  padding: 12px;
  color: #000000;
  background-color: #ffffff;
  cursor: pointer;
}

.dropdown-item:hover {
  background-color: #f1f1f1;
}

#status {
  text-align: center;
  min-height: 20px;
}

/* Themes */
.theme-light {
  background-color: #ffffff;
  color: #000000;
}

.theme-light .dropdown-btn {
  background-color: #d3b3aa;
  color: #000000;
}

.theme-dark {
  background-color: #121212;
  color: #ffffff;
}

.theme-dark .dropdown-btn {
  background-color: #4f4f4f;
  color: #ffffff;
}

.theme-ocean {
  background-color: #0077be;
  color: #e0f7fa;
}

.theme-ocean .dropdown-btn {
  background-color: #193e51;
  color: #ffffff;
}

.theme-nord {
  background-color: #2e3440;
  color: #d8dee9;
}

.theme-nord .dropdown-btn {
  background-color: #596680;
  color: #ffffff;
}
```

```js
const themes = [
  {
    name: "light",
    message: "Hello sunshine — Light theme is on!"
  },
  {
    name: "dark",
    message: "The night is yours — Dark theme is on!"
  },
  {
    name: "nord",
    message: "The frost has settled - Nord theme is on!"
  },
  {
    name: "ocean",
    message: "Blue skies and high tides — Ocean theme is on!"
  },
];

const setTheme = (theme) => {
  document.body.className = "";
  document.body.classList.add(`theme-${theme.name}`);
};

const updateStatus = (theme) => {
  const status = document.getElementById("status");
  status.textContent = theme.message;
};

document.addEventListener("DOMContentLoaded", () => {
  themes.forEach((theme) => {
    const button = document.getElementById(`theme-${theme.name}`);
    
    if (button) {
      button.addEventListener("click", () => {
        setTheme(theme);
        updateStatus(theme);
        closeMenu();
      });
    }
  });

  const toggleButton = document.getElementById("theme-switcher-button");
  const menu = document.getElementById("theme-dropdown");

  toggleButton.addEventListener("click", () => {
    const isExpanded = toggleButton.getAttribute("aria-expanded") === "true";
    if (isExpanded) {
      closeMenu();
    } else {
      openMenu();
    }
  });

  function closeMenu() {
    menu.setAttribute("hidden", "");
    toggleButton.setAttribute("aria-expanded", "false");
  }

  function openMenu() {
    menu.removeAttribute("hidden");
    toggleButton.setAttribute("aria-expanded", "true");
  }

  // Close the menu when clicking outside
  document.addEventListener("click", (e) => {
    if (!toggleButton.contains(e.target) && !menu.contains(e.target)) {
      closeMenu();
    }
  });
});
```
