---
id: 65f03d9f92eac9183a4d3281
title: Hatua ya 2
challengeType: 20
dashedName: step-2
---

# --description--

Vector inaweza kufafanuliwa kwa kuratibu mbili, `x` na `y`, katika ndege ya Euclidean. Umbali kati ya chanzo cha mihimili na nukta `(x, y)` utakuwa urefu wake, au norm. Na mwelekeo wa vector utaelekezwa kuelekea `(x, y)`.

<img alt="vector ya vipimo viwili (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Ndani ya darasa la vitu `Vector`, tengeneza njia `__init__` na uipe vigezo vitatu, `self`, `x`, na `y`.

# --hints--

Unapaswa kufafanua njia `__init__` ndani ya darasa la vitu `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

Njia yako ya `__init__` inapaswa kuchukua vigezo vitatu: `self`, `x`, na `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
