---
id: 65f01969115f933073b6be03
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Vector ni kitu chenye urefu (au ukubwa) na mwelekeo na hakiwezi kuwakilishwa na nambari moja tu. Katika fizikia, vector hutumika sana kuwakilisha nguvu, kasi, kuongezeka kwa kasi, na kiasi kingine chochote.

Katika mradi huu utajenga nafasi ya vector, seti ambayo mfululizo wa operesheni umefafanuliwa kati ya vipengele vilivyomo katika seti hiyo. Utajifunza maelezo yote hivi karibuni.

Kwa sasa, anza mradi kwa kutangaza darasa la vitu tupu linaloitwa `Vector`.

# --hints--

Unapaswa kutangaza darasa la vitu tupu linaloitwa `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
