---
id: 65f0694fb2296f3caadf8347
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Urefu wa vekta $\mathbf{a}$, au norm, kawaida huonyeshwa kama $\\| \mathbf{a} \\|$. Unaweza kuhesabiwa kama mzizi wa mraba wa jumla ya sehemu zake zilizopimwa mraba:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

Hesabu norm ya vekta na rudisha matokeo kutoka kwa njia yako ya `norm`.

# --hints--

Njia yako ya `norm` inapaswa kutumia fomula iliyotolewa kuhesabu na kurudisha nambari inayowakilisha norm ya mfano wa vekta uliopo. Usikaribishe thamani.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
