---
id: 65f058e06f34fd1f0ee6e55d
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Ifuatayo, ndani ya darasa la `Vector`, tengeneza njia tupu ya `norm` na umpe kigezo cha `self`.

# --hints--

Darasa lako la `Vector` linapaswa kuwa na njia ya `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

Njia yako ya `norm` inapaswa kuwa na kigezo cha `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
