---
id: 65f07c9b1ffb814d856dcffc
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Jaribu kwamba `norm` inafanya kazi kama inavyotarajiwa kwa kuchapisha thamani inayorejeshwa na njia hiyo.

# --hints--

Unapaswa kuita `norm()` kwenye `v1` na kuchapisha matokeo.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
