---
id: 65f055b9190fc41ca35549b8
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Python hutoa njia mbalimbali ambazo zina alama mbili za underscore mwanzoni na mwishoni mwa majina yao. Huenda tayari unafahamu baadhi, kama `__init__` na `__str__`. Njia hizi, zinazofuata muundo wa jina la `__<name>__`, huitwa njia maalum, njia za kichawi, au njia za *dunder* (zinazoashiria double underscore).

Kufafanua njia maalum katika darasa la vitu huathiri tabia ya darasa hilo. Huitwa kwa siri katika hali maalum (mfano `__init__` wakati wa kuunda mfano, `__str__` wakati kitu kinachapishwa au kupelekwa kwa `str()`). Katika mradi huu, utajifunza baadhi ya njia maalum zinazotumika sana.

Kwa sasa, weka thamani ya `x` kwa sifa ya kipengele `x` ya kitu cha `Vector`.

# --hints--

Unapaswa kuweka thamani ya `x` kwa `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
