---
id: 65f06a8e5a57673d700c79c3
title: Hatua 7
challengeType: 20
dashedName: step-7
---

# --description--

Nje ya darasa la `Vector`, tengeneza mfano wa `Vector` ukipitisha nambari kamili `2` na `3` kama thamani za `x` na `y` na uweke thamani yake katika kigezo kinachoitwa `v1`.

# --hints--

Unapaswa kutangaza kigezo `v1` na kuwekea thamani `Vector(2, 3)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
