---
id: 65f3fa60a93b84110b3f1708
title: Hatua 10
challengeType: 20
dashedName: step-10
---

# --description--

Unataka kurudisha usemi wa mfuatano wa herufi wenye maana. Kwa mfano, katika kesi ya `v1`, unataka kurudisha mfuatano wa herufi unaoonyesha tuple yenye thamani za sehemu za vector: `(2, 3)`.

Kutoka kwa njia ya `__str__`, rudisha mfuatano wa herufi unaowakilisha vector kama tuple yenye sehemu za vector kwa mpangilio. Kisha, toka nje ya darasa la `Vector` na chapisha `v1` ili ukague matokeo.

# --hints--

Njia yako ya `__str__` inapaswa kurudisha mfuatano wa herufi unaowakilisha vector kama tuple yenye sehemu za vector kwa mpangilio.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

Unapaswa kuchapisha `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
