---
id: 65f3fe07cc763212efe91285
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Vector inaweza kuwa na nambari `n` ya vipimo (sehemu). Hapa kuna uwakilishi wa vector yenye vipimo 3:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Hadi sasa, umeunda vector yenye vipimo 2. Unataka uwezekano wa kuwakilisha vector zenye nambari tofauti ya vipimo bila kuandika upya msimbo unaohitajika kwa kila kesi maalum. Kwa hiyo, utatumia urithi.

Anza kwa kubadilisha jina la darasa la `Vector` kuwa `R2Vector` ili kubainisha kuwa darasa hili litashughulikia vector zenye vipimo 2. Kumbuka pia kubadilisha uanzishaji wa `v1`.

# --hints--

Unapaswa kubadilisha jina la darasa la `Vector` kuwa `R2Vector`. Kumbuka pia kubadilisha uanzishaji wa `v1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
