---
id: 65f412c208c3791fee305acf
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Katika Python, unaweza kulazimisha matumizi ya hoja za maneno tu kwa kuongeza `*` kama hoja ya ziada katika saini ya kitendakazi au njia. Badilisha njia zote za `__init__` kwa kuongeza `*` kama kigezo cha pili (baada ya `self`). Kigezo chochote kilichowekwa baada ya hicho kitahitaji matumizi ya hoja ya neno katika wito wa kitendakazi/njia.

Hii inamaanisha kwamba unahitaji pia kubadilisha wito wa `super().__init__(x, y)`. Fanya hivyo kwa kumpa `x` thamani ya `x`, na `y` thamani ya `y`.

Mwishowe, badilisha uanzishaji wa `v1` na `v2` kwa kutumia hoja za maneno.

# --hints--

Unapaswa kulazimisha hoja za maneno kwa kuongeza `*` kama kigezo cha pili katika njia zako zote za `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

Unapaswa kubadilisha wito wako wa `super.__init__(x, y)` ili kutumia hoja za maneno.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

Unapaswa kubadilisha mgawo wa thamani wa `v1` kuwa `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

Unapaswa kubadilisha mgawo wa thamani wa `v2` kuwa `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
