---
id: 65f40051d6b09a139f253e8e
title: Hatua 12
challengeType: 20
dashedName: step-12
---

# --description--

Urithi unakuwezesha kufafanua darasa la vitu kutoka kwa lililopo tayari. Darasa jipya, linaloitwa mtoto, hurithi njia zote na vigezo vya darasa lililopo, linaloitwa mzazi.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

Katika mfano ulio juu, darasa la mtoto `Oak` hurithi kutoka `Tree` na hurithi njia ya `sprout` kutoka kwa darasa la mzazi `Tree`.

Tengeneza darasa jipya linaloitwa `R3Vector` na fuata mfano ulio juu ili lifanye urithi kutoka kwa darasa la `R2Vector`.

# --hints--

Unapaswa kufafanua darasa jipya `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

Darasa lako jipya `R3Vector` linapaswa kurithi kutoka `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
