---
id: 65f8709620b2ce1a62608f5a
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Unapohitaji kupata sifa fulani kwa njia ya mabadiliko kuanzia kwenye ingizo la mfuatano wa herufi, kitendakazi iliyo ndani ya `getattr()` ndicho unachohitaji. Kinachukua kitu kama hoja yake ya kwanza, na mfuatano wa herufi unaoonyesha jina la sifa ya kipengele kama sifa yake ya pili.

Anza kurekebisha njia ya `__str__` kwa kubadilisha mfuatano wa herufi unaorejeshwa na `__str__()` na usemi wa kizalishi unaorudia kupitia sifa za kipengele za kitu na kuita `getattr()` kwa kila sifa `i`.

# --hints--

Unapaswa kurejesha usemi wa kizalishi unaorudia juu ya `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

Unapaswa kurejesha usemi wa kizalishi unaotumia `i` kama kigezo cha kurudia.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

Unapaswa kurejesha usemi wa kizalishi unaoitisha `getattr(self, i)` kwa kila kipengele `i` kilichopo katika `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
