---
id: 65f9bc10a9fb1612e066e9e8
title: Hatua 42
challengeType: 20
dashedName: step-42
---

# --description--

Unapoongeza vekta mbili, kila sehemu ya vekta moja huongezwa na sehemu sawa ya vekta nyingine. Kwa mfano, kuongeza `(1, 2)` na `(2, 4)` huunda vekta ya tatu `(3, 6)`, ambapo `3` ni jumla ya `1` na `2` na `6` ni jumla ya `2` na `4`.

Kamusi ya `kwargs` itakuwa na jozi za ufunguo-thamani zinazohitajika kuanzisha vekta mpya ya darasa la vitu sawa na vekta mbili zilizoongezwa pamoja.

Badilisha kamusi tupu kuwa utengenezaji wa kamusi unaorudia kupitia `vars(self)` na kwa kila ufunguo (`i`) unda jozi ya ufunguo-thamani, ambapo ufunguo ni ufunguo sawa wa kurudia kwa sasa na thamani yake ni jumla ya thamani za sifa ya kipengele `i` za wapatanishi wawili.

# --hints--

Unapaswa kubadilisha kamusi tupu imepewa thamani `kwargs` kuwa utengenezaji wa kamusi unaorudia juu ya `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

Utengenezaji wa kamusi umepewa thamani `kwargs` unapaswa kutumia kigezo `i` kurudia juu ya `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

Utengenezaji wa kamusi umepewa thamani `kwargs` unapaswa kutumia `i` kama ufunguo.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

Utengenezaji wa kamusi umepewa thamani `kwargs` unapaswa kuweka jumla ya sifa ya kipengele `i` ya `self` na sifa ya kipengele `i` ya `other` kwa ufunguo `i` kwa kila `i` katika `vars(self)`. Tumia kitendakazi `getattr()` kwa hilo.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
