---
id: 65f9cb678070ca1668898c70
title: Hatua 44
challengeType: 20
dashedName: step-44
---

# --description--

Kwa njia ile ile `__add__` inavyofanywa wito chini ya kifuniko wakati vitu viwili vinaongezwa pamoja, njia ya `__sub__` hufanywa wito kwa njia isiyo ya moja kwa moja katika hali ya kutoa.

Sasa, fafanua njia tupu ya `__sub__` na toa vigezo viwili: `self`, na `other`. Ndani ya njia yako mpya, tengeneza taarifa ya `if` ili ukague kama `self` na `other` haviko katika darasa la vitu moja na rudisha `NotImplemented`, kama ulivyofanya awali.

# --hints--

Unapaswa kufafanua njia ya `__sub__` ndani ya darasa la vitu la `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

Njia yako ya `__sub__` inapaswa kuwa na vigezo viwili, `self` na `other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

Unapaswa kuunda taarifa ya `if` inayokagua kama `self` na `other` haviko katika darasa la vitu moja.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

Unapaswa kurudisha `NotImplemented` kutoka ndani ya mwili wa `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
