---
id: 7a97b79c096ecb000d03bd36
title: Jenga kalkuleta ya nambari ya Nth Fibonacci
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `fibonacci`.
2. Unapaswa kufafanua safu ya data inayoitwa `sequence` ndani ya kitendakazi cha `fibonacci`, na inapaswa kuanzishwa na thamani za `[0, 1]`.
3. Kitendakazi cha `fibonacci` kinapaswa kupokea kigezo kimoja, nambari kamili isiyo hasi `n`.
4. Kufanya wito wa `fibonacci(n)` kunapaswa kutumia mbinu ya programu inayobadilika (dynamic programming) kuhesabu na kurudisha nambari ya `n`-th kutoka katika mfululizo wa Fibonacci, ambapo kila nambari ni jumla ya nambari mbili zilizotangulia.
5. Kila nambari ya Fibonacci iliyohesabiwa inapaswa kuongezwa kwenye safu ya data ya `sequence`.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `fibonacci`.

```js
assert.isFunction(fibonacci);
```

Kitendakazi chako cha `fibonacci` kinapaswa kuchukua kigezo kimoja.

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

Unapaswa kuwa na safu ya data inayoitwa `sequence` ndani ya kitendakazi cha `fibonacci` iliyowekwa na `[0, 1]`.

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` inapaswa kurudisha `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` inapaswa kurudisha `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` inapaswa kurudisha `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` inapaswa kurudisha `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` inapaswa kurudisha `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` inapaswa kurudisha `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` inapaswa kurudisha `610`.

```js
assert.equal(fibonacci(15), 610);
```

Hauruhusiwi kutumia kurudia (recursion) katika msimbo wako.

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
