---
id: 56533eb9ac21ba0edf2244b3
title: 建置一個攝氏轉華氏轉換器
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

在本實驗中，你將撰寫一個函式（程式），用來將溫度從攝氏轉換為華氏。將攝氏轉換為華氏的公式是：

```js
fahrenheit = celsius * (9/5) + 32
```

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個名為 `convertCtoF` 的函式（程式）。
2. `convertCtoF` 應該接受一個單一的數值引數，該引數是攝氏溫度。
3. `convertCtoF` 應該傳回一個數字。

# --hints--

你應該創建一個名為 `convertCtoF` 的函式（程式）。

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` 應該接受一個單一參數。

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` 應該傳回一個數字。

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` 應該傳回一個值為 `-22`。

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` 應該傳回一個值為 `14`。

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` 應該傳回一個值為 `32`。

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` 應該傳回一個值為 `68`。

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` 應該傳回一個值為 `86`。

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
