---
id: 56533eb9ac21ba0edf2244b3
title: 創建一個攝氏度轉華氏度的轉換器
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

在本實驗中，你將編寫一個將攝氏溫度轉換爲華氏溫度的函數。 將攝氏度換算成華氏度的公式是：

```js
fahrenheit = celsius * (9/5) + 32
```

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶故事：**

1. 你應創建一個名爲 `convertCtoF` 的函數。
2. `convertCtoF` should take a single numeric argument, which is the temperature in Celsius.
3. `convertCtoF` 應返回一個數字。

# --hints--

你應創建一個名爲 `convertCtoF` 的函數。

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` 應只接受一個參數。

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` should return a number.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` should return a value of `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` should return a value of `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` should return a value of `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` should return a value of `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` should return a value of `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
