---
id: 684ccbdc1ef89e042b744e09
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

الجزء التالي من ورشة العمل هو إضافة زر الإغلاق.

أسفل عنصر `p` الخاص بك، أضف عنصر `button` مع خاصية `aria-label` مضبوطة على `"Dismiss"`. داخل عنصر `button` هذا، أضف عنصر `span` مع خاصية `aria-hidden`. يجب أن يكون نص عنصر `span` هو `&times;`.

*ملاحظة*: بما أن هذه الورشة ستركز فقط على العمل مع أنماط Tailwind CSS، فلن يُضاف أي تفاعل إلى زر `button`.

# --hints--

يجب أن يكون لديك عنصر `button`.

```js
assert.exists(document.querySelector("button"));
```
 
يجب أن يحتوي عنصر `button` الخاص بك على خاصية `aria-label`.

```js
const buttonEl = document.querySelector("button");
assert.isTrue(buttonEl?.hasAttribute("aria-label"));
```

يجب أن تحتوي خاصية `aria-label` في عنصر `button` الخاص بك على القيمة `"Dismiss"`.

```js
const buttonEl = document.querySelector("button");
assert.equal(buttonEl?.getAttribute("aria-label"), "Dismiss");
```

يجب أن يحتوي عنصر `button` الخاص بك على عنصر `span` مسكن بداخله.

```js
assert.exists(document.querySelector("button span"));
```

يجب أن يحتوي عنصر `span` الخاص بك على خاصية `aria-hidden`.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl?.hasAttribute("aria-hidden"));
```

يجب أن يحتوي عنصر `span` الخاص بك على النص `&times;`.

```js
const spanEl = document.querySelector("span");
assert.equal(spanEl?.innerHTML, "×");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>
  </body>
</html>
```
