---
id: 684cff0bf0a4ca0ad7e42990
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

الخطوة التالية هي تدوير زوايا الحدود. لتعيين خاصية `border-radius` باستخدام Tailwind CSS، يمكنك استخدام صنف مساعد بالنمط التالي: `rounded-size`، حيث تمثل `size` نصف القطر.

إليك مثالًا لتطبيق انحناء زوايا كبير جدًا على عنصر:

```html
<p class="rounded-xl">This is border radius.</p>
```

داخل وسم `div` الافتتاحي، طبق انحناء زوايا بحجم `md`.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على الصنف المساعد `rounded-md`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
