---
id: 684cc59eed2706fa50e8933c
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتعلم كيف تبدأ العمل بأساسيات Tailwind CSS من خلال بناء مكوِّن رسالة خطأ.

في الخطوات الأولى، ركز على بناء ترميز HTML الأولي.

داخل عنصر `body`، أضف عنصر `div`. ثم داخل ذلك العنصر `div`، أضف عنصر `p`.

# --hints--

يجب أن يكون لديك عنصر `div`.

```js
assert.exists(document.querySelector("div"));
```

يجب أن يكون لديك عنصر `p`.

```js
assert.exists(document.querySelector("p"));
```

يجب أن يكون عنصر `p` داخل عنصر `div`.

```js
assert.exists(document.querySelector("div p"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body>
</html>
```
