---
id: 684cca09aac0fc035db4099f
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

داخل عنصر `p`، أضف النص `Error! Something went wrong. Please try again.` بحيث يكون `Error!` ملفوفًا داخل عنصر `strong`.

# --hints--

يجب أن يحتوي عنصر `p` على عنصر `strong` بداخله.

```js
assert.exists(document.querySelector("p strong"));
```

يجب أن يحتوي عنصر `strong` على النص `Error!`.

```js
assert.equal(document.querySelector("strong")?.textContent, "Error!");
```

يجب أن يحتوي عنصر `p` على النص الصحيح الذي ينتهي بـ `Something went wrong. Please try again.`

```js
const remainingText = "Something went wrong. Please try again."
const pElement = document.querySelector("p");
assert.isTrue(pElement?.textContent.trim().endsWith(remainingText));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <p>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </p>
    </div>
  </body>
</html>
```
