---
id: 684cd8bf5445de089af9c837
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

حان الوقت الآن لتطبيق لون حدود. لتعيين خاصية `border-color` باستخدام Tailwind CSS، يمكنك استخدام صنف مساعد بالنمط التالي: `border-color-number`، حيث يمثل `number` الدرجة اللونية للحدود.

إليك مثالًا لتطبيق لون حدود بنفسجي على عنصر:

```html
<p class="border-purple-400">This is border-color.</p>
```

داخل وسم `div` الافتتاحي، طبق لون حدود أحمر بدرجة `300`.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على الصنف المساعد `border-red-300`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("border-red-300"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
