---
id: 684cd027744c2a0520ba4d58
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

يستخدم Tailwind CSS أصناف الأدوات التي تُطبَّق مباشرة على الترميز. تشمل هذه الأصناف تنسيق الخطوط، الألوان، تخطيطات Flexbox، الشبكة والمزيد.

لتطبيق لون خلفية في Tailwind CSS، استخدم البادئة `bg`، تليها اسم اللون ورقم. الرقم يُمثل درجة اللون.

إليك مثالًا لتطبيق لون خلفية أزرق داكن على عنصر فقرة.

```html
<p class="bg-blue-900">This has a dark blue background.</p>
```

في وسم `div` الافتتاحي، أضف صنفًا يحتوي على صنف الأداة `bg-red-100` الذي يُمثل لون خلفية أحمر فاتح.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على خاصية صنف.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.hasAttribute("class"));
```

يجب أن تحتوي خاصية الصنف على صنف الأداة `bg-red-100`.

```js
const divEl = document.querySelector("div");
assert.exists(document.querySelector("div.bg-red-100"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
