---
id: 684d160a2370b01310ae0fda
title: الخطوة 18
challengeType: 0
dashedName: step-18
---

# --description--

اللون الأحمر الجديد يبدو جميلًا، لكن سيكون أفضل لو كان النص أكبر. لتعيين `font-size` لعنصر باستخدام Tailwind CSS، يمكنك استخدام النمط التالي: `text-value`.

إليك مثالًا على تعيين `font-size` لعنصر إلى `0.875rem` أو (`14px`):

```html
<p class="text-sm">This is for font-size.</p>
```

داخل وسم `p` الافتتاحي، استخدم صنف الأداة الصحيح لتعيين `font-size` إلى `xl` الذي يمثل `1.25rem` أو (`20px`).

# --hints--

يجب أن يحتوي وسم `p` الافتتاحي على صنف الأداة `text-xl`.

```js
const pEl = document.querySelector("p");
assert.isTrue(pEl.classList.contains("text-xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      --fcc-editable-region--
      <p class="text-red-700">
      --fcc-editable-region--
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
