---
id: 684d0eeb9c81e5109e52fe8b
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

عنصر `div` بدأ يبدو أفضل. لكن الآن زر الإغلاق `button` قريب جدًا من عنصر الفقرة.

لإنشاء مساحة بين عناصر المرونة، يمكنك استخدام النمط التالي: `gap-number`، حيث تمثل كل وحدة `number` مقدار 0.25rem (أو 4px) من المساحة الثابتة بين الصفوف والأعمدة في تخطيط المرونة.

إليك مثالًا على تطبيق `gap` على عنصر:

```html
<p class="flex gap-8">This is gap.</p>
```

داخل وسم `div` الافتتاحي، استخدم صنف الأداة المناسب لتعيين `gap` إلى `16px`.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأداة `gap-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("gap-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
