---
id: 684d17a3b350bb13d3c19c0f
title: الخطوة 19
challengeType: 0
dashedName: step-19
---

# --description--

حان الوقت الآن لتطبيق بعض الأنماط على عنصر `button`.

ابدأ بإضافة خاصية صنف إلى عنصر `button` الخاص بك. بعد ذلك، طبق صنف الأدوات المناسب لتعيين لون نص أحمر بدرجة `700`.

ثم، داخل وسم `button` الافتتاحي، استخدم صنف الأدوات الصحيح لتعيين `font-size` إلى `2xl` الذي يمثل `1.5rem` أو (`24px`).

# --hints--

يجب أن يحتوي عنصر `button` على خاصية صنف.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.hasAttribute("class"));
```

يجب أن يحتوي وسم `button` الافتتاحي على صنف الأدوات `text-red-700`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl?.classList.contains("text-red-700"));
```

يجب أن يحتوي وسم `button` الافتتاحي على صنف الأدوات `text-2xl`.

```js
const btnEl = document.querySelector("button");
assert.isTrue(btnEl.classList.contains("text-2xl"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4 justify-center">
      <p class="text-red-700 text-xl">
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      --fcc-editable-region--
      <button aria-label="Dismiss">
      --fcc-editable-region--
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
