---
id: 684d009994a3840b98c9aa3a
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

عرض عنصر `div` الآن يمتد على كامل عرض الصفحة. لتطبيق العروض في Tailwind CSS، لديك عدة خيارات، مثل استخدام نمط `w-number` للعروض الثابتة أو نمط `w-fraction` للنسب المئوية.

إليك مثالًا لتطبيق عرض محدد بـ `66%` من الحاوية الأصلية:

```html
<p class="w-2/3">This is a percentage width.</p>
```

داخل وسم `div` الافتتاحي، طبّق عرضًا محددًا بـ `50%` من الحاوية الأصلية.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأدوات `w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
