---
id: 684d09cad4ec900e41db2e74
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

الآن، عنصر `div` موضوع في الزاوية العلوية اليسرى من شاشة المعاينة. لكن من الأفضل إضافة بعض الفراغ في أعلى العنصر. لتطبيق `margin` على عنصر باستخدام Tailwind CSS، يمكنك استخدام النمط التالي: `m-number`، حيث تمثل كل وحدة `number` مقدار 0.25rem (أو 4 بكسل) لتطبيقها على جميع الجوانب الأربعة.

إذا أردت تطبيق الهامش على جانب واحد فقط، يمكنك استخدام أي من الأنماط التالية: `mt-number` (هامش أعلى)، `mb-number` (هامش أسفل)، `mr-number` (هامش يمين) و `ml-number` (هامش يسار).

إليك مثالًا لتطبيق `margin-bottom` بمقدار 32px على عنصر:

```html
<p class="mb-8">This is margin.</p>
```

داخل وسم `div` الافتتاحي، استخدم صنف الأدوات الصحيح لتطبيق `margin-top` بمقدار `16px`.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأدوات `mt-4`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("mt-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
