---
id: 684d105e488dbd1166a717ae
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

إذا زدت حجم نافذة المعاينة، ستلاحظ أن كل العناصر لا تزال مصطفة إلى اليسار داخل الحاوية. سيكون من الجيد أن تُوسَط عناصر المرن.

لوضع عناصر المرن على طول المحور الرئيسي للحاوية باستخدام Tailwind CSS، يمكنك استخدام النمط التالي: `justify-value`. يمكن أن تكون `value` مثل `start`، `end`، `center`، `between` والمزيد. يمكن استخدام أي قيمة صالحة في CSS العادي مع `justify-content` هنا مع Tailwind CSS.

داخل وسم `div` الافتتاحي، استخدم صنف الأدوات الصحيح لضبط خاصية `justify-content` على `center`.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأدوات `justify-center`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("justify-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>  
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 md:mx-auto flex gap-4">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
