---
id: 684d03df2a0ddf0c735865f0
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

الآن ترى أن عرض العنصر هو نصف حجم الـ `body`. هذا يعمل جيدًا للأجهزة الأكبر لكنه لا يعمل بنفس الكفاءة للأجهزة الأصغر.

يتبع Tailwind CSS نهج التصميم الموجه للأجهزة المحمولة أولًا. لذا إذا أردت إنشاء تصميمات متجاوبة لأجهزة أخرى، يمكنك استخدام المتغيرات المتجاوبة في Tailwind مثل `md` و `lg` بهذه الطريقة:

```html
<p class="lg:w-2/3">This is a percentage width.</p>
```

المثال التالي سيضبط العرض ليكون `66%` من العنصر الأب لأحجام الأجهزة التي تساوي `1024px` أو أكثر.

داخل وسم `div` الافتتاحي، عدل `w-1/2` ليُطبق فقط على أحجام الأجهزة التي تساوي `768px` أو أكثر. يمكنك إضافة بادئة `md` لهذه الصنف.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأدوات `md:w-1/2`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:w-1/2"));
```

يجب ألا يحتوي وسم `div` الافتتاحي على الصنف `w-1/2` بعد الآن.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md w-1/2">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
