---
id: 68501d32b69f9d7756334ffc
title: الخطوة 13
challengeType: 0
dashedName: step-13
---

# --description--

إذا كبرت نافذة المعاينة، يجب أن ترى عنصر `div` في وسط الصفحة أفقيًا. لكن على الأجهزة الأصغر، يمتد عرض عنصر `div` ليشمل عرض الصفحة بالكامل.

حدّث صنف `mx-auto` ليُطبّق فقط على أجهزة `md` وما فوق.

# --hints--

يجب أن يحتوي وسم `div` الافتتاحي على صنف الأدوات `md:mx-auto`.

```js
const divEl = document.querySelector("div");
assert.isTrue(divEl.classList.contains("md:mx-auto"));
```

يجب ألا يحتوي وسم `div` الافتتاحي على صنف `mx-auto` بعد الآن.

```js
const divEl = document.querySelector("div");
assert.isFalse(divEl.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Error Message Component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div class="bg-red-100 border-2 border-red-300 rounded-md md:w-1/2 p-4 mt-4 mx-auto">
    --fcc-editable-region--
      <p>
        <strong>Error!</strong> Something went wrong. Please try again.
      </p>
      <button aria-label="Dismiss">
        <span aria-hidden>&times;</span>
      </button>
    </div>
  </body>
</html>
```
