---
id: a5deed1811a43193f9f1c841
title: 實作 Element Skipper
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

在本實驗中，你將創建一個函式（程式），該函式會跳過陣列中的元素，直到根據特定的測試函式（程式）找到可接受的元素為止。

例如，對於像 `[1, 1, 1, 2, 1, 1, 1]` 這樣的陣列和測試函式 `function(n) {return n === 2}`，第一個符合條件的元素是在索引 3，因此所有在此之前的元素都需要被捨棄，輸出應該是剩下的元素 `[2, 1, 1, 1]`。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事**

1. 你應該有一個 `dropElements` 函式（程式），它接受一個陣列（`arr`）和一個函式（程式）（`func`）作為引數。
1. `dropElements` 函式（程式）應該迭代陣列並從第一個元素開始移除元素，直到 `func` 對某個元素傳回 `true`。
1. 如果條件符合，`dropElements` 函式應該傳回陣列中剩餘的元素。
1. 如果條件從未被滿足，則應傳回空陣列。

# --hints--

你應該有一個 `dropElements` 函式（程式）。

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` 應該傳回 `[3, 4]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` 應該傳回 `[1, 0, 1]`。

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` 應該傳回 `[1, 2, 3]`。

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` 應該傳回 `[]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` 應該傳回 `[7, 4]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` 應該傳回 `[3, 9, 2]`。

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
