---
id: 66ad9c633fa26d3c1475eae3
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

現在，是時候為機器人添加另一則訊息。

為你的程式碼添加另一個 `console` 述語，以記錄訊息 `"Allow me to introduce myself."`。

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

你的程式碼中應該有第三個 `console.log()` 述語。

```js
assert.lengthOf(logSpy.calls, 3);
```

你的 `console` 述語應該有訊息 `"Allow me to introduce myself."`。不要忘記引號。

```js
assert.deepInclude(logSpy.calls, ["Allow me to introduce myself."]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
