---
id: 66adc5919c1853448f119d11
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

現在看起來機器人想要更改他們的暱稱。

使用重新指派，將字串 `"awesomeTeacherBot"` 指定給 `bot` 變數。

# --hints--

你應該將字串 `"awesomeTeacherBot"` 指定給 `bot` 變數。

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
