---
id: 66adb844118ba74107ce771f
title: 步驟 8
challengeType: 1
dashedName: step-8
---

# --description--

機器人的下一條信息將涉及機器人的位置。

創建一個名爲 `botLocationSentence` 的變量。

然後使用 `+` 運算符進行字符串連接，連接字符串 `"I live in "` 和 `botLocation` 變量，之後是一個句號（`.`）。

將此值賦值給 `botLocationSentence` 變量。

然後，將 `botLocationSentence` 的值輸出到控制檯。

# --hints--

你應該有一個 `botLocationSentence` 變量。

```js
assert.isNotNull(botLocationSentence);
```

你的 `botLocationSentence` 變量應存儲一個字符串的值。

```js
assert.isString(botLocationSentence);
```

你應該使用 `+` 運算符進行字符串連接，連接字符串 `"I live in "` 和 `botLocation` 變量，之後是一個句號（`.`）。 注意空格。

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

你的代碼中應包含一個 `console` 語句。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

你的 `console` 語句應輸出 `botLocationSentence` 變量。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
