---
id: 66adc42868cab843ccee87d9
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

現在是時候查看新的 `bot` 值。

首先建立一個名為 `nicknameIntroduction` 的新變數。

使用字串串接將字串 `"My nickname is "` 與 `bot` 變數連接，後面接一個句號（`.`）。

將結果字串指定到 `nicknameIntroduction` 變數。

然後，將 `nicknameIntroduction` 的值記錄到主控台。 

# --hints--

你應該有一個名為 `nicknameIntroduction` 的變數。

```js
assert.isNotNull(nicknameIntroduction);
```

你的 `nicknameIntroduction` 變數應該是一個字串。

```js
assert.isString(nicknameIntroduction);
```

你應該使用字串串接將字串 `"My nickname is "` 與 `bot` 變數連接，後面接一個句點（`.`）。請注意空格。

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

你應該在你的程式碼中有一個 `console` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

你應該將 `nicknameIntroduction` 的值記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
