---
id: 66ad8d150264db3926eccfeb
title: 步驟 3
challengeType: 1
dashedName: step-3
---

# --description--

在之前的課程中，你學習了 `let` 關鍵字以及如何在 JavaScript 中宣告變數。 

這裡提醒你如何使用 `let` 關鍵字宣告變數：

```js
let greeting;
```

使用 `let` 關鍵字宣告一個名為 `bot` 的變數。

**注意**：你在這裡使用 `let` 是因為在工作坊的後面，你會更改 `bot` 變數的值。

# --hints--

你應該在你的程式碼中有 `let` 關鍵字。

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

你應該宣告一個名為 `bot` 的變數。

```js
assert.isNotNull(bot);
```

你應該使用 `let` 關鍵字來宣告變數。

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
