---
id: 66adcf383276814776aba3ca
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

接著，創建一個名為 `favoriteSubjectSentence` 的變數。

使用字串串接將字串 `"My favorite subject is "` 與 `favoriteSubject` 變數連接，後面加上句號。

將結果指定到 `favoriteSubjectSentence` 變數。

然後，將 `favoriteSubjectSentence` 的值記錄到主控台。

# --hints--

你應該有一個名為 `favoriteSubjectSentence` 的變數。

```js
assert.isNotNull(favoriteSubjectSentence);
```

你的 `favoriteSubjectSentence` 變數應該是一個字串。

```js
assert.isString(favoriteSubjectSentence);
```

你應該使用字串串接將字串 `"My favorite subject is "` 與 `favoriteSubject` 變數連接，後面加上句號。請注意空格。

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

你應該在你的程式碼中有一個 `console` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

你應該將 `favoriteSubjectSentence` 的值記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
