---
id: 66ad984e92ce213a66f7e51c
title: 步驟4
challengeType: 1
dashedName: step-4
---

# --description--

當你需要用多個單詞聲明變量時，可以使用 <dfn>camelCase</dfn> 命名法。

使用 `camelCase` 時，第一個單詞全部小寫，隨後每個單詞的第一個字母大寫。

這是一個例子：

```js
let thisIsCamelCase;
```

使用 `let` 聲明一個名爲 `botLocation` 的變量。

# --hints--

你應該使用 `let` 來聲明變量。

```js
assert.lengthOf(code.match(/let/g), 2);
```

你的變量應命名爲 `botLocation` 。

```js
assert.isNotNull(botLocation);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
