---
id: 66ada3f46945763dd97f43f8
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

之前，你已經創建了 `bot` 和 `botLocation` 變數。現在，你將使用它們將新的訊息輸出到主控台。

在之前的課程中，你學會了如何使用 `+` 運算子來串接字串，如下所示：

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

請記得在將字串與變數串接時要注意空格。

創建一個名為 `botIntroduction` 的變數。

然後使用字串串接與 `+` 運算子將字串 `"My name is "`、接著是 `bot` 變數，再接著是一個句號（`.`）串接起來。 

將此值指定到 `botIntroduction` 變數。

然後，將 `botIntroduction` 變數記錄到主控台。

# --hints--

你應該有一個名為 `botIntroduction` 的變數。

```js
assert.isNotNull(botIntroduction);
```

你的 `botIntroduction` 變數應該保存一個字串的值。

```js
assert.isString(botIntroduction);
```

你應該使用字串串接和 `+` 運算子將字串 `"My name is "` 與 `bot` 變數連接，後面接一個句點（`.`）。請注意空格。將此串接後的字串指定給 `botIntroduction` 變數。

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

你的程式碼中應該有第四個 `console.log()` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

你的 `console` 述語應該輸出 `botIntroduction` 變數。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
