---
id: 66adbc8a9793d64250f7e609
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

在之前的課程中，你學會了如何像這樣重新指定變數的值：

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

使用重新指派，將字串 `"professorBot"` 指派給 `bot` 變數。

# --hints--

你應該將字串 `"professorBot"` 指定給 `bot` 變數。

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
