---
id: 66ad9c633fa26d3c1475eae3
title: 步驟 6
challengeType: 1
dashedName: step-6
---

# --description--

現在，是時候為機器人添加另一則訊息。 

為你的程式碼添加另一個 `console` 述語，以記錄訊息 `"Allow me to introduce myself."`。

# --hints--

你的程式碼中應該有第三個 `console.log()` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

你的 `console` 述語應該有訊息 `"Allow me to introduce myself."`。不要忘記引號。

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
