---
id: 66ad8294a0ad902f1b31b612
title: 步驟 1
challengeType: 1
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個問候機器人來學習如何使用 JavaScript 基礎知識。 

在這第一個步驟中，你會想從 greeting bot 輸出一則訊息到主控台。

請記住你在之前的課程中學過 `console.log()` 和字串。 

這裡提醒你如何使用 `console.log()` 搭配字串：

```js
console.log("Hello, World!");
```

為你的程式碼添加一個 `console.log()` 述語，將字串 `"Hi there!"` 輸出到主控台。不要忘記在訊息周圍加上引號！ 


# --hints--

你應該在你的程式碼中有一個 `console.log()` 述語。

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

你的 `console` 述語應該有訊息 `"Hi there!"`。別忘了引號。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
