---
id: 66ad8ab945d266383f318cbf
title: 步驟 2
challengeType: 1
dashedName: step-2
---

# --description--

現在你應該會在主控台看到機器人的第一則訊息。 

是時候為機器人添加第二則訊息。

為你的程式碼添加另一個 `console.log` 述語，以將訊息 `"I am excited to talk to you."` 輸出到主控台。

# --hints--

你的程式碼中應該有第二個 `console.log()` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

你的 `console` 述語應該包含訊息 `"I am excited to talk to you."`。不要忘記引號。

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
