---
id: 66adb844118ba74107ce771f
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

機器人接下來的訊息將關於機器人的位置。

創建一個名為 `botLocationSentence` 的變數。

然後使用字串串接與 `+` 運算子將字串 `"I live in "` 與 `botLocation` 變數連接，接著加上句號（`.`）。

將此值指定給 `botLocationSentence` 變數。

然後，將 `botLocationSentence` 的值記錄到主控台。

# --hints--

你應該有一個 `botLocationSentence` 變數。

```js
assert.isNotNull(botLocationSentence);
```

你的 `botLocationSentence` 變數應該保存一個字串的值。

```js
assert.isString(botLocationSentence);
```

你應該使用字串串接和 `+` 運算子，將字串 `"I live in "` 與 `botLocation` 變數連接，後面接上一個句點（`.`）。請注意空格。

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

你應該在你的程式碼中有一個 `console` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

你的 `console` 述語應該輸出 `botLocationSentence` 變數。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
