---
id: 66adc6046acc18453decf577
title: 步驟 12
challengeType: 1
dashedName: step-12
---

# --description--

要看到機器人的新暱稱，你需要將新的訊息記錄到主控台。

創建一個名為 `newNicknameGreeting` 的新變數。

然後使用字串串接將字串 `"I love my nickname but I wish people would call me "` 與 `bot` 變數連接，接著加上一個句號。

將結果指定到 `newNicknameGreeting` 變數。

然後，將 `newNicknameGreeting` 的值記錄到主控台。

# --hints--

你應該有一個名為 `newNicknameGreeting` 的變數。

```js
assert.isNotNull(newNicknameGreeting);
```

你的 `newNicknameGreeting` 變數應該是一個字串。

```js
assert.isString(newNicknameGreeting);
```

你應該使用字串串接將字串 `"I love my nickname but I wish people would call me "` 與 `bot` 變數連接，後面接上一個句號。

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

你應該在你的程式碼中有一個 `console` 述語。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

你應該將 `newNicknameGreeting` 的值記錄到主控台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
