---
id: 66ad998285df023b23dacdd3
title: 步驟 5
challengeType: 1
dashedName: step-5
---

# --description--

現在，是時候初始化 `botLocation` 變數。

當你需要宣告多字的變數時，你可以使用 <dfn>camelCase</dfn> 命名慣例。

使用 `camelCase` 時，第一個字全部小寫，後續每個字的第一個字母皆大寫。 

這裡有一個範例：

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

使用 `let` 關鍵字在同一行宣告並指定字串 `"the universe"` 給 `botLocation` 變數。

# --hints--

使用 `let` 關鍵字宣告變數。

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

你的變數應該命名為 `botLocation`。

```js
assert.isNotNull(botLocation);
```

你應該為你的 `botLocation` 變數指定一個字串。

```js
assert.isString(botLocation);
```

你的 `botLocation` 變數應該被指定為字串 `"the universe"`。

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
