---
id: 66adc82170acb1464bd348eb
title: 步驟 13
challengeType: 1
dashedName: step-13
---

# --description--

機器人最後幾則訊息將專注於機器人最喜歡的主題。

首先建立一個名為 `favoriteSubject` 的新變數，並將字串 `"Computer Science"` 指定給它。

# --hints--

你應該有一個名為 `favoriteSubject` 的變數。

```js
assert.isNotNull(favoriteSubject);
```

你的 `favoriteSubject` 變數應該是一個字串。

```js
assert.isString(favoriteSubject);
```

你應該將字串 `"Computer Science"` 指定給 `favoriteSubject` 變數。

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
