---
id: 66ad984e92ce213a66f7e51c
title: 步驟 4
challengeType: 1
dashedName: step-4
---

# --description--

現在，是時候為你的 `bot` 變數指定值。

在之前的課程中，你學會了如何像這樣為變數指定值：

```js
variableName = "Here is the value";
```

請記住，指定運算子 `=` 右側的內容是你要指定給左側變數的值。

將字串 `"teacherBot"` 指定到 `bot` 變數。

# --hints--

你應該為你的 `bot` 變數指定一個字串。

```js
assert.isString(bot);
```

你的 `bot` 變數應該被指定為字串 `"teacherBot"`。

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
