---
id: 691f7bbf76229bb97c60827d
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви потренуєтеся використовувати цикли `for...of`, створюючи функцію, яка рахує, скільки разів рядок зустрічається в масиві рядків.

Перед тим як написати цю функцію, ви створите простішу, яка виводить у консоль кожен символ рядка.

Почніть із визначення пустої функції з назвою `printCharacters` з параметром `str`.

# --hints--

Ви маєте створити функцію `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Ваша функція `printCharacters` повинна мати один параметр з назвою `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
