---
id: 6920d189e36b6b60a6330fee
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер, всередині циклу, виведіть у консоль змінну `char`.

# --hints--

У вас має бути інструкція `console.log` з аргументом `char` всередині циклу.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Коли функція `printCharacters` викликається з аргументом `'test'`, інструкція `console.log()` всередині функції має викликатися чотири рази з такими аргументами, у порядку: `'t'`, `'e'`, `'s'` і `'t'`. Натомість вона була викликана з такими аргументами: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
