---
id: 6920c9cda1c7f056f75e4b1a
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер настав час написати свій перший цикл `for...of`. У попередніх уроках ви дізналися, що можна використовувати цикл `for...of` для ітерування над значеннями з ітерованого об’єкта (наприклад, рядків або масивів).

Ось приклад циклу `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Додайте цикл `for...of` з порожнім блоком коду всередині вашої функції. Він має ітерувати над кожним символом у аргументі `str`.

# --hints--

Ваша функція `printCharacters` має містити цикл `for...of` всередині свого блоку коду.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Ваш цикл `for...of` має ітерувати над аргументом `str` вашої функції.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
