---
id: 6921d688b8d58581d34459b9
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

У наступних кількох кроках ви створите функцію, яка рахує, скільки разів рядок зустрічається в масиві рядків.

Для початку визначте порожню функцію з іменем `getMatchedWordCount` з параметрами `sentence` та `match` у такому порядку.

# --hints--

Ви маєте створити функцію `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Ваша функція `getMatchedWordCount` має мати два параметри: `sentence` та `match`, у такому порядку.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
