---
id: 6921dae10700e300f26994e3
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Щоб побачити вихідні дані вашої функції `getMatchedWordCount` у її поточному стані, додайте інструкцію `console.log` нижче. Усередині `console.log()` викличте функцію з аргументом `["I", "really", "really", "really", "like", "to", "code"]` для параметра `sentence` та `"really"` для параметра `match`.

# --hints--

Ви повинні викликати функцію `getMatchedWordCount` всередині інструкції `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Ви повинні викликати функцію `getMatchedWordCount` з аргументом `["I", "really", "really", "really", "like", "to", "code"]` для параметра `sentence` та `"really"` для параметра `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
