---
id: 6921d688b8d58581d34459b9
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

У наступних кількох кроках ви створите функцію, яка рахує, скільки разів рядок зустрічається в масиві рядків.

Для початку визначте порожню функцію з іменем `getMatchedWordCount` з параметрами `sentence` та `match` у такому порядку.

# --hints--

Ви маєте створити функцію `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Ваша функція `getMatchedWordCount` має мати два параметри: `sentence` та `match`, у такому порядку.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
