---
id: 6921e10ae2681205b516fbe6
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Тепер час додати логіку до вашого циклу, яка відповідно збільшує `count`.

Всередині циклу використайте умовну інструкцію, щоб збільшити `count` на `1`, якщо змінна `word` дорівнює змінній `match`. Інакше залиште значення `count` без змін.

# --hints--

Виклик `getMatchedWordCount([], "bar")`, `getMatchedWordCount(["foo"], "bar")`, `getMatchedWordCount(["foo", "bar", "foo"], "bar")` і `getMatchedWordCount(["foo", "bar", "foo"], "foo")` у такому порядку має повертати значення типу `"number"`; натомість він повернув значення типу `--fcc-actual--`.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

`getMatchedWordCount([], "bar")` має повертати `0`, але повертає `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

`getMatchedWordCount(["foo"], "bar")` має повертати `0`, але повертає `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

`getMatchedWordCount(["foo", "bar", "foo"], "bar")` має повертати `1`, але повертає `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

`getMatchedWordCount(["foo", "bar", "foo"], "foo")` має повертати `2`, але повертає `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
