---
id: 6921d90b5ce50185505d2004
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Оскільки функція `getMatchedWordCount` має повертати числовий лічильник, почніть із створення змінної-фішки.

Створіть змінну з іменем `count` за допомогою ключового слова `let`, ініціалізуйте її значенням `0` і додайте інструкцію `return`, яка повертає змінну `count`.

# --hints--

Ви повинні оголосити змінну-фішку `count` всередині функції `getMatchedWordCount`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

Щоб змінну-фішку `count` можна було пізніше збільшувати, оголосіть її з ключовим словом `let`.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

У вашій функції має бути присвоєння значення `0` змінній-фішці `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

У вашій функції має бути інструкція `return`, яка повертає змінну-фішку `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

Ваша функція `getMatchedWordCount` має повертати значення `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
