---
id: 6921de8d2a111d0341cff702
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Чудово, ваша функція повертає `0`, але поки що нічого не рахує.

Щоб це виправити, всередині `getMatchedWordCount` створіть цикл `for...of` з порожнім блоком коду, який ітерує над кожним словом у `sentence`.

# --hints--

Ваша функція `getMatchedWordCount` має містити цикл `for...of` всередині свого блоку коду.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Ваш цикл `for...of` має ітерувати над аргументом `sentence` вашої функції.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
