---
id: 587d7fb4367417b2b2512c00
title: Erweitern Sie Ihr Projekt mit externen Paketen von npm
challengeType: 2
forumTopicId: 301527
dashedName: expand-your-project-with-external-packages-from-npm
---

# --description--

Einer der größten Vorteile eines Paketmanagers ist seine leistungsstarke Abhängigkeitsverwaltung. Anstatt manuell sicherstellen zu müssen, dass Sie alle Abhängigkeiten erhalten, wenn Sie ein Projekt auf einem neuen Computer einrichten, installiert npm automatisch alles für Sie. Aber wie kann npm genau wissen, was Ihr Projekt benötigt? Lernen Sie den `dependencies`-Abschnitt Ihrer package.json-Datei kennen.

In diesem Abschnitt werden die Pakete gespeichert, die Ihr Projekt benötigt, und zwar im folgenden Format:

```json
"dependencies": {
  "package-name": "version",
  "express": "4.14.0"
}

```

# --instructions--

Fügen Sie die Version `1.1.0` des Pakets `@freecodecamp/example` zum `dependencies`-Feld Ihrer `package.json`-Datei hinzu.

**Hinweis:** `@freecodecamp/example` ist ein fiktives Paket, das als Lernwerkzeug verwendet wird.

# --hints--

`"dependencies"` sollte `"@freecodecamp/example"` enthalten.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Die Version von `"@freecodecamp/example"` sollte `"1.1.0"` sein.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^[\^\~]?1\.1\.0/,
        'Wrong version of "@freecodecamp/example" installed. It should be 1.1.0'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
