---
id: 587d7fb5367417b2b2512c03
title: "Verwenden Sie das Caret-Zeichen, um die neueste Minor-Version einer Abhängigkeit zu nutzen"
challengeType: 2
forumTopicId: 301531
dashedName: use-the-caret-character-to-use-the-latest-minor-version-of-a-dependency
---

# --description--

Ähnlich wie die Tilde, die wir in der letzten Aufgabe kennengelernt haben und die npm erlaubt, den neuesten PATCH einer Abhängigkeit zu installieren, erlaubt das Caret-Zeichen (`^`) npm auch zukünftige Updates zu installieren. Der Unterschied besteht darin, dass das Caret sowohl MINOR-Updates als auch PATCHes zulässt.

Ihre aktuelle Version von `@freecodecamp/example` sollte `~1.2.13` sein, was npm erlaubt, auf die neueste `1.2.x`-Version zu aktualisieren. Wenn Sie stattdessen das Caret-Zeichen (^) als Präfix für die Version verwenden, würde npm erlaubt sein, auf jede `1.x.x`-Version zu aktualisieren.

```json
"package": "^1.3.8"
```

Dies würde Updates auf jede `1.x.x`-Version des Pakets erlauben.

# --instructions--

Verwenden Sie das Caret-Zeichen (`^`) als Präfix für die Version von `@freecodecamp/example` in Ihren Abhängigkeiten, damit npm es auf jede neue MINOR-Version aktualisieren kann.

**Hinweis:** Die Versionsnummern selbst sollten nicht geändert werden.

# --hints--

`"dependencies"` sollte das Caret-Zeichen (`"@freecodecamp/example"`) enthalten.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Die `"@freecodecamp/example"`-Version sollte mit `"^1.x.x"` übereinstimmen.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\^1\./,
        'Wrong version of "@freecodecamp/example". It should be ^1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
