---
id: 587d7fb5367417b2b2512c02
title: "Verwenden Sie das Tilde-Zeichen, um immer die neueste Patch-Version einer Abhängigkeit zu nutzen"
challengeType: 2
forumTopicId: 301532
dashedName: use-the-tilde-character-to-always-use-the-latest-patch-version-of-a-dependency
---

# --description--

In der letzten Aufgabe haben Sie npm angewiesen, nur eine bestimmte Version eines Pakets zu verwenden. Das ist eine nützliche Methode, um Ihre Abhängigkeiten einzufrieren, wenn Sie sicherstellen müssen, dass verschiedene Teile Ihres Projekts miteinander kompatibel bleiben. In den meisten Fällen möchten Sie jedoch keine Bugfixes verpassen, da diese oft wichtige Sicherheitsupdates enthalten und (hoffentlich) dabei keine Fehler verursachen.

Um einer npm-Abhängigkeit zu erlauben, auf die neueste PATCH-Version zu aktualisieren, können Sie die Version der Abhängigkeit mit dem Tilde-Zeichen (`~`) als Präfix versehen. Hier ist ein Beispiel, wie Sie Updates auf jede `1.3.x`-Version erlauben.

```json
"package": "~1.3.8"
```

# --instructions--

In der Datei package.json lautet Ihre aktuelle Regel, wie npm `@freecodecamp/example` aktualisieren darf, dass eine bestimmte Version (`1.2.13`) verwendet wird. Jetzt möchten Sie jedoch die neueste `1.2.x`-Version erlauben.

Verwenden Sie das Tilde-Zeichen (`~`) als Präfix für die Version von `@freecodecamp/example` in Ihren Abhängigkeiten, damit npm Updates auf jede neue _Patch_-Version durchführen darf.

**Hinweis:** Die Versionsnummern selbst sollten nicht geändert werden.

# --hints--

`"dependencies"` sollte `"@freecodecamp/example"` enthalten.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.property(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" does not include "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

`"@freecodecamp/example"` Version sollte mit `"~1.2.13"` übereinstimmen.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.match(
        data.dependencies['@freecodecamp/example'],
        /^\~1\.2\.13/,
        'Wrong version of "@freecodecamp/example". It should be ~1.2.13'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
