---
id: 587d7fb4367417b2b2512bfe
title: Fügen Sie eine Lizenz zu Ihrer package.json hinzu
challengeType: 2
forumTopicId: 301523
dashedName: add-a-license-to-your-package-json
---

# --description--

Das `license`-Feld ist der Ort, an dem Sie Benutzer darüber informieren, was sie mit Ihrem Projekt tun dürfen.

Einige gängige Lizenzen für Open-Source-Projekte sind MIT und BSD. Lizenzinformationen sind nicht zwingend erforderlich, und Urheberrechtsgesetze in den meisten Ländern geben Ihnen standardmäßig die Eigentümerschaft an dem, was Sie erstellen. Es ist jedoch immer eine gute Praxis, ausdrücklich anzugeben, was Benutzer dürfen und was nicht. Hier ist ein Beispiel für das Lizenzfeld:

```json
"license": "MIT",
```

# --instructions--

Füllen Sie das `license`-Feld in der package.json-Datei Ihres Projekts nach eigenem Ermessen aus.

# --hints--

package.json sollte einen gültigen "license"-Schlüssel enthalten

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.license, '"license" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
