---
id: 587d7fb4367417b2b2512bfd
title: Schlüsselwörter zu Ihrer package.json hinzufügen
challengeType: 2
forumTopicId: 301526
dashedName: add-keywords-to-your-package-json
---

# --description--

Das Feld `keywords` ist der Ort, an dem Sie Ihr Projekt mit passenden Schlüsselwörtern beschreiben können. Hier ist ein Beispiel:

```json
"keywords": [ "descriptive", "related", "words" ],
```

Wie Sie sehen, ist dieses Feld als Array von doppelt angeführten Strings aufgebaut.

# --instructions--

Fügen Sie dem Feld `keywords` in der package.json-Datei Ihres Projekts ein Array mit passenden Strings hinzu.

Eines der Schlüsselwörter sollte "freecodecamp" sein.

# --hints--

package.json sollte einen gültigen Schlüssel "keywords" enthalten

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.keywords, '"keywords" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

Das Feld "keywords" sollte ein Array sein

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.isArray(data.keywords, '"keywords" is not an array');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```

"keywords" sollte "freecodecamp" enthalten

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.include(
        data.keywords,
        'freecodecamp',
        '"keywords" does not include "freecodecamp"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
