---
id: 587d7fb4367417b2b2512bff
title: Eine Version zu Ihrer package.json hinzufügen
challengeType: 2
forumTopicId: 301525
dashedName: add-a-version-to-your-package-json
---

# --description--

Eine `version` ist eines der Pflichtfelder Ihrer package.json-Datei. Dieses Feld beschreibt die aktuelle Version Ihres Projekts. Hier ist ein Beispiel:

```json
"version": "1.2.0",
```

# --instructions--

Fügen Sie der package.json-Datei Ihres Projekts eine `version` hinzu.

# --hints--

package.json sollte einen gültigen "version"-Schlüssel enthalten

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.version, '"version" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
