---
id: 587d7fb3367417b2b2512bfc
title: Eine Beschreibung zu Ihrer package.json hinzufügen
challengeType: 2
forumTopicId: 301522
dashedName: add-a-description-to-your-package-json
---

# --description--

Der nächste wichtige Teil einer guten package.json-Datei ist das `description`-Feld; hier gehört eine kurze, aber aussagekräftige Beschreibung Ihres Projekts hinein.

Falls Sie irgendwann planen, ein Paket bei npm zu veröffentlichen, ist dies der String, der Ihre Idee dem Benutzer verkaufen soll, wenn er entscheidet, ob er Ihr Paket installieren möchte oder nicht. Das ist jedoch nicht der einzige Anwendungsfall für die Beschreibung, sie ist eine großartige Möglichkeit, zusammenzufassen, was ein Projekt macht. Sie ist in jedem Node.js-Projekt ebenso wichtig, um anderen Entwicklerinnen und Entwicklern, zukünftigen Wartenden oder sogar Ihrem zukünftigen Ich zu helfen, das Projekt schnell zu verstehen.

Unabhängig davon, was Sie mit Ihrem Projekt vorhaben, ist eine Beschreibung auf jeden Fall empfehlenswert. Hier ist ein Beispiel:

```json
"description": "A project that does something awesome",
```

# --instructions--

Fügen Sie eine `description` zur package.json-Datei Ihres Projekts hinzu.

**Hinweis:** Denken Sie daran, für Feldnamen doppelte Anführungszeichen (") und Kommas (,) zur Trennung der Felder zu verwenden.

# --hints--

package.json sollte einen gültigen "description"-Schlüssel enthalten

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.exists(data.description, '"description" is missing');
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
