---
id: 587d7fb5367417b2b2512c04
title: Ein Paket aus Ihren Abhängigkeiten entfernen
challengeType: 2
forumTopicId: 301530
dashedName: remove-a-package-from-your-dependencies
---

# --description--

Sie haben nun einige Möglichkeiten getestet, wie Sie die Abhängigkeiten Ihres Projekts über den dependencies-Abschnitt der package.json verwalten können. Sie haben auch externe Pakete hinzugefügt, indem Sie sie in die Datei aufgenommen haben, und npm sogar mitgeteilt, welche Arten von Versionen Sie wünschen, indem Sie Sonderzeichen wie die Tilde oder das Caret-Zeichen verwendet haben.

Aber was, wenn Sie ein externes Paket entfernen möchten, das Sie nicht mehr benötigen? Sie haben es vielleicht schon erraten: Entfernen Sie einfach das entsprechende Schlüssel-Wert-Paar für dieses Paket aus Ihren Abhängigkeiten.

Diese Methode gilt auch für das Entfernen anderer Felder in Ihrer package.json.

# --instructions--

Entfernen Sie das `@freecodecamp/example`-Paket aus Ihren Abhängigkeiten.

**Hinweis:** Achten Sie darauf, nach dem Entfernen die richtige Anzahl an Kommas zu haben.

# --hints--

`"dependencies"` sollte `"@freecodecamp/example"` nicht enthalten.

```js
fetch(code + '/_api/package.json')
  .then(response => response.json())
  .then(
    data => {
      assert.notProperty(
        data.dependencies,
        '@freecodecamp/example',
        '"dependencies" still includes "@freecodecamp/example"'
      );
    },
    error => {
      throw new Error(error.message || error.responseText);
    }
  );
```
