---
id: 66ba762af611169359d9d369
title: Eine Zeitungsartikel erstellen
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten das Wurzel-`font-size` Ihres HTML-Dokuments auf `24px` setzen.
1. Sie sollten ein Element mit einer Klasse `newspaper` haben, das alle anderen Elemente enthält.
1. Ihr `.newspaper`-Element sollte eine `font-size` von `16px` und eine Schriftart von `Open Sans` mit einer Fallback-Schriftart von `sans-serif` haben.
1. Innerhalb Ihres `.newspaper`-Elements sollten Sie mindestens sieben weitere Elemente haben: eines für den Zeitungsnamen mit der Klasse `name`, eines für das Datum des Artikels mit der Klasse `date`, eines für die Überschrift mit der Klasse `headline`, eines für die Unterüberschrift mit der Klasse `sub-headline`, eines für den Autor mit der Klasse `author` und zwei Absätze, jeweils mit der Klasse `text`. Alle diese Elemente sollten mit Ihren Artikeldaten gefüllt sein.
1. Ihr `.name`-Element sollte eine `font-size` haben, die doppelt so groß ist wie die `font-size` des Wurzelelements, und die Schriftart `Times New Roman` verwenden, mit einer Fallback-Schriftart von `serif`.
1. Ihre `.name`- und `.author`-Elemente sollten per CSS alle ihre Zeichen in Großbuchstaben darstellen.
1. Ihr `.headline`-Element sollte eine `font-size` haben, die doppelt so groß ist wie die `font-size` seines Elternelements, und fett sein.
1. Ihr `.sub-headline`-Element sollte eine `font-weight` von `100`, eine `font-size`, die `1.5`-mal so groß ist wie die `font-size` seines Elternelements, und kursiv sein.
1. Ihr `.author` sollte per CSS fett dargestellt werden.
1. Ihre `.text`-Elemente sollten eine `text-indent` von `20px` haben.
1. Ihre `.text`-Elemente sollten eine `line-height` haben, die doppelt so groß ist wie die `font-size` ihres Elternelements.
1. Der erste Buchstabe Ihrer `.text`-Elemente sollte fett und doppelt so groß sein wie die `font-size` ihres Elternelements. Verwenden Sie dafür den `::first-letter`-Selektor.

# --hints--

Sie sollten einen `html`-Selektor haben, der seine `font-size` auf `24px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

Sie sollten ein Element mit der Klasse `newspaper` haben.

```js
assert.exists(document.querySelector('.newspaper'));
```

Sie sollten einen `.newspaper`-Selektor haben, der die `font-size` seiner Elemente auf `16px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

Sie sollten einen `.newspaper`-Selektor haben, der die `font-family` seiner Elemente auf `'Open Sans', sans-serif` setzt.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

Sie sollten ein Element mit der Klasse `name` innerhalb Ihres `.newspaper`-Elements haben.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

Ihr `.name`-Element sollte den Namen Ihrer Zeitung enthalten.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

Sie sollten ein Element mit der Klasse `date` innerhalb Ihres `.newspaper`-Elements haben.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

Ihr `.date`-Element sollte das Datum enthalten.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

Sie sollten ein Element mit der Klasse `headline` innerhalb Ihres `.newspaper`-Elements haben.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

Ihr `.headline`-Element sollte Ihre Artikelüberschrift enthalten.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

Sie sollten ein Element mit der Klasse `sub-headline` innerhalb Ihres `.newspaper`-Elements haben.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

Ihr `.sub-headline`-Element sollte Ihre Artikelunterüberschrift enthalten.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

Sie sollten ein Element mit der Klasse `author` innerhalb Ihres `.newspaper`-Elements haben.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

Ihr `.author`-Element sollte den Autor Ihres Artikels enthalten.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

Sie sollten mindestens zwei Absatz-Elemente haben, jeweils mit der Klasse `text`, innerhalb Ihres `.newspaper`-Elements.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

Ihre `.text`-Elemente sollten den Text Ihres Artikels enthalten.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

Sie sollten einen `.name`-Selektor haben, der die `font-size` seiner Elemente auf `2rem` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

Sie sollten einen `.name`-Selektor haben, der die `font-family` seiner Elemente auf `'Times New Roman', serif` setzt.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

Das `.name`-Element sollte die `text-transform`-Eigenschaft auf `uppercase` gesetzt haben.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Das `.author`-Element sollte die `text-transform`-Eigenschaft auf `uppercase` gesetzt haben.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

Sie sollten einen `.headline`-Selektor haben, der die `font-size` seiner Elemente auf `2em` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

Sie sollten einen `.headline`-Selektor haben, der die `font-weight` seiner Elemente auf `bold` setzt.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

Sie sollten einen `.sub-headline`-Selektor haben, der die `font-weight` seiner Elemente auf `100` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

Sie sollten einen `.sub-headline`-Selektor haben, der die `font-size` seiner Elemente auf `1.5em` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

Sie sollten einen `.sub-headline`-Selektor haben, der die `font-style` seiner Elemente auf `italic` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

Sie sollten einen `.author`-Selektor haben, der die `font-weight` seiner Elemente auf `bold` setzt.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

Sie sollten einen `.text`-Selektor haben, der die `text-indent` seiner Elemente auf `20px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

Sie sollten einen `.text`-Selektor haben, der die `line-height` seiner Elemente auf `2em` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

Sie sollten einen `.text::first-letter`-Selektor haben, der die `font-weight` seiner Elemente auf `bold` setzt.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

Sie sollten einen `.text::first-letter`-Selektor haben, der die `font-size` seiner Elemente auf `2em` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
