---
id: 5dc1798ff86c76b9248c6eb3
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Vipengele vya kichwa cha sehemu `h1` hadi `h6` vinatumika kuonyesha umuhimu wa maudhui yaliyo chini yao. Nambari ndogo inaonyesha umuhimu mkubwa zaidi, kwa hivyo vipengele vya `h2` vina umuhimu mdogo kuliko vipengele vya `h1`. 

```html
<h1>most important heading element</h1>
<h2>second most important heading element</h2>
<h3>third most important heading element</h3>
<h4>fourth most important heading element</h4>
<h5>fifth most important heading element</h5>
<h6>least important heading element</h6>
```

Tumia kipengele kimoja tu cha `h1` kwa kila ukurasa na weka vichwa vya sehemu vyenye umuhimu mdogo chini ya vichwa vya sehemu vyenye umuhimu mkubwa.

Chini ya kipengele cha `h1`, ongeza kipengele cha `h2` chenye maandishi haya:

`Cat Photos`

# --hints--

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/h1\>/));
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`. Ondoa ziada.

```js
assert(
  document.querySelector('h1') && document.querySelectorAll('h1').length === 1
);
```

Maandishi ya kipengele chako cha `h1` yanapaswa kuwa `CatPhotoApp`. Umeacha maandishi au kuna makosa ya tahajia.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

Lebos za `h2` zinapaswa kuwa kwa herufi ndogo. Kwa mkataba, lebo zote za HTML huandikwa kwa herufi ndogo.

```js
assert.notMatch(code, /<\/?H2>/);
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('h2'));
```

Kipengele chako cha `h2` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/h2\>/));
```

Maandishi ya kipengele chako cha `h2` yanapaswa kuwa `Cat Photos`. Weka maandishi `Cat Photos` tu kati ya lebo za kufungua na kufunga za `h2`.

```js
assert(document.querySelector('h2').innerText.toLowerCase() === 'cat photos');
```

Kipengele chako cha `h2` kinapaswa kuwa chini ya kipengele cha `h1`. Kipengele cha `h1` kina umuhimu mkubwa zaidi na kinapaswa kuwa juu ya kipengele cha `h2`.

```js
const collection = [...document.querySelectorAll('h1,h2')].map(
  (node) => node.nodeName
);
assert(collection.indexOf('H1') < collection.indexOf('H2'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>CatPhotoApp</h1>
    
--fcc-editable-region--
  </body>
</html>
```
