---
id: 5dc174fcf86c76b9248c6eb2
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Vipengele vya HTML vina lebo za kufungua kama `<h1>` na lebo za kufunga kama `</h1>`. Maandishi ambayo kipengele kitaonyesha huwekwa kati ya lebo zake za kufungua na kufunga.

Badilisha maandishi ya kipengele cha `h1` hapa chini kutoka `Hello World` hadi `CatPhotoApp` na uangalie mabadiliko kwenye onyesho la kivinjari.

Ukipomaliza, bonyeza kitufe cha "Check Your Code" kuona kama ni sahihi.

# --hints--

Maandishi `CatPhotoApp` yanapaswa kuwepo kwenye msimbo. Labda unataka kukagua tahajia yako.

```js
assert(code.match(/catphotoapp/i));
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('h1'));
```

Lebos za `h1` zinapaswa kuwa kwa herufi ndogo. Kwa mkataba, lebo zote za HTML huandikwa kwa herufi ndogo.

```js
assert.notMatch(code, /<\/?H1>/);
```

Kipengele chako cha `h1` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/h1\>/));
```

Una zaidi ya kipengele kimoja cha `h1`. Ondoa kipengele cha ziada cha `h1`.

```js
assert(document.querySelectorAll('h1').length === 1);
```

Maandishi ya kipengele chako cha `h1` yanapaswa kuwa `CatPhotoApp`. Huenda umeacha maandishi, kuna makosa ya tahajia, au hayako kati ya lebo za kufungua na kufunga za kipengele cha `h1`.

```js
assert(document.querySelector('h1').innerText.toLowerCase() === 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--
    <h1>Hello World</h1>
--fcc-editable-region--
  </body>
</html>
```
