---
id: 5dc24165f86c76b9248c6ebc
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Vipengele vyote vya `img` vinapaswa kuwa na sifa ya kipengele ya `alt`. Maandishi ya sifa ya kipengele ya `alt` hutumiwa na kisomaji cha skrini kuboresha ufikikaji na huonyeshwa ikiwa picha haitapakia.

Hapa kuna mfano wa kipengele cha `img` chenye sifa ya kipengele ya `alt`:

```html
<img src="cat.jpg" alt="A cat">
```

Ndani ya kipengele cha `img`, ongeza sifa ya kipengele ya `alt` yenye maandishi haya:

`A cute orange cat lying on its back`

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `img`. Umetoa kipengele cha `img` kutoka hatua ya awali.

```js
assert(document.querySelector('img'));
```

Kipengele chako cha `img` hakina sifa ya kipengele ya `alt`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert(document.querySelector('img').hasAttribute('alt'));
```

Thamani ya sifa ya kipengele ya `alt` katika kipengele cha `img` imewekwa kuwa tofauti na `A cute orange cat lying on its back`. Hakikisha thamani ya sifa ya kipengele ya `alt` imezungukwa na alama za nukuu.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert(altText.match(/A cute orange cat lying on its back\.?$/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```
