---
id: 5ddbd81294d8ddc1510a8e56
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Maandishi ya kiungo lazima yawe kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha nanga (`a`).

Hapa kuna mfano wa kiungo chenye maandishi `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Ongeza maandishi ya nanga `cat photos` kwenye kipengele cha nanga. Hii itakuwa maandishi ya kiungo.

# --hints--

Kipengele chako cha nanga (`a`) kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('a'));
```

Kipengele chako cha nanga (`a`) kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/a\>/));
```

Maandishi ya kipengele chako cha nanga (`a`) yanapaswa kuwa `cat photos`. Hakikisha unaweka maandishi ya kiungo kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha nanga (`a`).

```js
assert.strictEqual(
  document.querySelector('a').innerText.toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
