---
id: 5dfa2407b521be39a3de7be1
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Ili kufungua viungo katika kichupo kipya, unaweza kutumia sifa ya kipengele `target` kwenye kipengele cha tagi ya nanga (`a`). 

Sifa ya kipengele `target` inaeleza wapi hati iliyounganishwa itafunguliwa. `target="_blank"` hufungua hati iliyounganishwa katika kichupo kipya au dirisha jipya.

Hapa kuna sintaksia ya msingi kwa kipengele cha `a` chenye sifa ya kipengele `target`:

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Ongeza sifa ya kipengele `target` yenye thamani `_blank` kwenye lebo ya kufungua ya kipengele cha tagi ya nanga (`a`), ili kiungo kifunguliwe katika kichupo kipya.

# --hints--

Kipengele chako cha `p` kinapaswa kuwa na kipengele cha tagi ya nanga (`a`) kilichopangwa ndani chenye maandishi `cat photos`. Huenda umekifuta au kuna makosa ya tahajia.

```js
const anchor = document.querySelectorAll('p > a');
assert.isNotEmpty(anchor);
assert.strictEqual(anchor[1]?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos')
```

Kipengele chako cha tagi ya nanga (`a`) hakina sifa ya kipengele `target`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert.isTrue(document.querySelectorAll('a')[1]?.hasAttribute('target'));
```

Thamani ya sifa ya kipengele `target` inapaswa kuwa `_blank`. Huenda umesahau thamani au kuna makosa ya tahajia. Kumbuka thamani za sifa za kipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert.strictEqual(document.querySelectorAll('a')[1]?.getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
