---
id: 5dfa37b9eacea3f48c6300b0
title: Hatua 22
challengeType: 0
dashedName: step-22
---

# --description--

Ili kuunda orodha isiyopangwa ya vitu, unaweza kutumia kipengele cha `ul`. 

Baada ya kipengele cha `h3` chenye maandishi ya `Things cats love:`, ongeza kipengele cha orodha isiyopangwa (`ul`). Kumbuka kwamba hakuna kitu kitakachoonyeshwa kwa sasa.

# --hints--

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('ul'));
```

Kipengele chako cha `ul` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/ul>/));
```

Kipengele cha `ul` kinapaswa kuwa juu ya lebo ya kufunga ya kipengele cha pili cha `section`.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')?.[1]?.lastElementChild;
assert.equal(secondSectionLastElemNode?.nodeName , 'UL');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
