---
id: 5ef9b03c81a63668521804df
title: Hatua 48
challengeType: 0
dashedName: step-48
---

# --description--

Sifa ya kipengele `id` hutumika kutambua vipengele maalum vya HTML. Kila thamani ya sifa ya kipengele `id` lazima iwe ya kipekee ikilinganishwa na thamani zote nyingine za `id` kwa ukurasa mzima.

Hapa kuna mfano wa kipengele cha `input` chenye sifa ya kipengele `id`:

```html
<input id="email">
```

Ongeza sifa ya kipengele `id` yenye thamani `indoor` kwa kitufe cha radio. Vipengele vinapokuwa na sifa nyingi, mpangilio wa sifa hizo hauna umuhimu.

# --hints--

Kitufe chako cha radio kinapaswa bado kuwa kati ya lebo ya kufungua na lebo ya kufunga ya kipengele cha `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')?.[0]?.childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

Kitufe chako cha radio kinapaswa kuwa na sifa ya kipengele `id`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('id'));
```

Kitufe chako cha radio kinapaswa kuwa na sifa ya kipengele `type`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('type'));
```

Kitufe chako cha radio kinapaswa kuwa na sifa ya kipengele `id` yenye thamani `indoor`. Umeacha thamani au umefanya makosa ya tahajia. Kumbuka kuwa thamani za sifa za kipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert.match(document.querySelector('input')?.id, /^indoor$/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input type="radio"> Indoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
