---
id: 5ef9b03c81a63668521804d8
title: Hatua 39
challengeType: 0
dashedName: step-39
---

# --description--

Kipengele cha `input` kinakuwezesha njia kadhaa za kukusanya data kutoka kwa fomu ya mtandao. Kama vipengele vya `img`, vipengele vya `input` ni <dfn>kipengele kisicho na maudhui</dfn> na havihitaji lebo za kufunga.

Panga ndani kipengele cha `input` ndani ya kipengele cha `form`.

# --hints--

Kipengele chako cha `form` kinapaswa kuwa na lebo ya kufungua na lebo ya kufunga kwa mpangilio sahihi. Huenda umekosa moja au zote mbili za lebo zinazohitajika, au umeziweka kwa mpangilio usio sahihi.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Leboa ya kufungua ya kipengele chako cha `form` inapaswa kuwa na sifa ya kipengele ya `action` pekee. Ondoa chochote kingine ulichokiandika ndani yake.

```js
assert([...document.querySelector('form').attributes].length < 2);
```

Unapaswa kuunda kipengele cha `input`. Kagua sintaksia.

```js
assert(document.querySelector('input'));
```

Kipengele chako cha `input` kinapaswa kuwa na lebo ya kufungua, lakini si lebo ya kufunga.

```js
assert(document.querySelector('input') && !code.match(/<\/input\>/g));
```

Kipengele chako cha `input` kinapaswa kupangwa ndani ya kipengele cha `form`.

```js
assert(document.querySelector('form > input'));
```

`form` inapaswa kuwa na kipengele cha `input` pekee ndani yake. Ondoa vipengele vyovyote vya HTML au maandishi kati ya lebo za kipengele cha `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
