---
id: 5ef9b03c81a63668521804d1
title: Hatua 28
challengeType: 0
dashedName: step-28
---

# --description--

Baada ya kipengele cha `figure`, ongeza kipengele kingine cha `h3` chenye maandishi:

`Top 3 things cats hate:`

# --hints--

Kipengele cha `h3` kinapaswa kuwepo juu ya lebo ya kufunga ya kipengele cha `section` cha pili. Hakikisha kina lebo ya kufungua na lebo ya kufunga.

```js
assert(
  document.querySelectorAll('main > section')[1].lastElementChild.nodeName ===
    'H3' && code.match(/<\/h3\>/g).length === 2
);
```

Kipengele kipya cha `h3` kinapaswa kuwa na maandishi ya `Top 3 things cats hate:`. Hakikisha unajumuisha nukta ya wima mwishoni mwa maandishi.

```js
assert(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase()
    .replace(/\s+/g, ' ') === 'top 3 things cats hate:'
);
```

Kipengele cha `figure` kinapaswa kuwepo juu ya kipengele kipya cha `h3`. Huenda ulifuta kipengele cha `figure` kwa bahati mbaya.

```js
const secondSectionLastElemNode = document.querySelectorAll('main > section')[1]
  .lastElementChild;
assert(
  secondSectionLastElemNode.nodeName === 'H3' &&
    secondSectionLastElemNode.previousElementSibling.nodeName === 'FIGURE'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
