---
id: 5ef9b03c81a63668521804e2
title: Hatua 56
challengeType: 0
dashedName: step-56
---

# --description--

Fomu hutumia mara nyingi visanduku vya kuchagua kwa maswali ambayo yanaweza kuwa na majibu zaidi ya moja. Kipengele cha `input` chenye sifa ya kipengele `type` iliyowekwa kuwa `checkbox` huunda kisanduku cha kuchagua.

Chini ya kipengele cha `legend` ulichokiweka, ongeza `input` yenye sifa ya kipengele `type` iliyowekwa kuwa `checkbox` na maandishi yake yakiwa:

` Loving`

# --hints--

Kipengele cha `input` kwa kisanduku chako cha kuchagua kinapaswa kuwa na lebo ya kufungua, lakini si lebo ya kufunga.

```js
assert.isTrue(document.querySelector('fieldset > input') && !code.match(/<\/input\>/g));
```

Unapaswa kuwa umeongeza kipengele kimoja tu cha ingizo kwa kisanduku chako cha kuchagua. Ondoa ziada zozote.

```js
assert.lengthOf(document.querySelectorAll('fieldset > input'), 1);
```

Kipengele chako kipya cha `input` kinapaswa kuwa chini ya kipengele cha `legend` chenye maandishi `What's your cat's personality?`. Umeviweka kwa mpangilio usio sahihi.

```js
const existingLegendElem = document.querySelectorAll('fieldset > legend')?.[1];
assert.strictEqual(
  existingLegendElem?.nextElementSibling?.nodeName,
  'INPUT'
);
```

Kipengele chako kipya cha `input` hakina sifa ya kipengele `type`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua.

```js
assert.isTrue(document.querySelector('fieldset > input')?.hasAttribute('type'));
```

Kipengele chako kipya cha `input` kinapaswa kuwa na sifa ya kipengele `type` yenye thamani `checkbox`. Umeacha thamani au umefanya makosa ya tahajia. Kumbuka thamani za sifa za kipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert.match(document.querySelector('fieldset > input')?.getAttribute('type'),/^checkbox$/i);
```

Ingawa umeweka sifa ya kipengele `input` ya kipengele kipya cha `type` kuwa `checkbox`, inashauriwa kila mara kuzuzua thamani ya sifa ya kipengele na alama za nukuu.

```js
assert(!/\<\s*input\s+type\s*=\s*checkbox/i.test(code));
```

Maandishi ` Loving` yanapaswa kuwekwa moja kwa moja upande wa kulia wa kisanduku chako cha kuchagua. Hakikisha kuna nafasi kati ya kipengele na maandishi. Umeacha maandishi au umefanya makosa ya tahajia.

```js
const checkboxInputElem = document.querySelector('input[type="checkbox"]');
assert.match(
  checkboxInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Loving/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
--fcc-editable-region--
            <legend>What's your cat's personality?</legend>
            
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
