---
id: 5ef9b03c81a63668521804dd
title: Hatua 47
challengeType: 0
dashedName: step-47
---

# --description--

Vipengele vya `label` hutumika kusaidia kuhusisha maandishi ya kipengele cha `input` na kipengele cha `input` yenyewe (hasa kwa teknolojia za msaada kama kisomaji cha skrini).

Hapa kuna mfano wa kipengele cha `label` chenye kitufe cha `radio`:

```html
<label><input type="radio"> cat</label>
```

Katika mfano huu, kubofya neno `"cat"` pia kutachagua kitufe cha `radio`.

Panga ndani kitufe chako cha `radio` ndani ya kipengele cha `label`.

# --hints--

Unapaswa kuhakikisha kitufe cha radio bado kipo.

```js
assert(document.querySelector('input[type="radio"]'));
```

Maandishi ya ` Indoor` yanapaswa kuwekwa moja kwa moja upande wa kulia wa kitufe chako cha `radio`. Hakikisha kuna nafasi kati ya kipengele na maandishi. Huenda umeacha maandishi au kuna makosa ya tahajia.

```js
const radioInputElem = document.querySelector('input');
assert.match(
  radioInputElem?.nextSibling?.nodeValue?.replace(/\s+/g, ' '),
  / Indoor/i
);
```

Kipengele chako cha `label` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('label'));
```

Kipengele chako cha `label` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/label\>/));
```

Kitufe chako cha radio na maandishi yake yanapaswa kuwekwa yote kati ya lebo za kufungua na kufunga za kipengele cha `label`.

```js
const labelChildNodes = [...document.querySelector('form > label').childNodes];
assert(
  labelChildNodes.filter((childNode) => childNode.nodeName === 'INPUT').length
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="radio"> Indoor
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
