---
id: 5ef9b03c81a63668521804db
title: Hatua 43
challengeType: 0
dashedName: step-43
---

# --description--

Ili kuzuia mtumizi kuwasilisha fomu yako wakati taarifa muhimu haipo, unahitaji kuongeza sifa ya kipengele `required` kwenye kipengele cha `input`.

Hapa kuna mfano wa sehemu ya ingizo yenye sifa ya kipengele `required`:

```html
<input type="text" name="firstName" required>
```

Haitaji kuweka thamani kwa sifa ya kipengele `required`. Badala yake, ongeza tu neno `required` kwenye kipengele cha `input`, ukihakikisha kuna nafasi kati yake na sifa nyingine.

# --hints--

Umefuta kipengele chako cha `input` au kina sintaksia isiyo sahihi. Thamani zote za sifa za kipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert(document.querySelectorAll('input').length);
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Futa vipengele vingine vya HTML au maandishi ndani ya kipengele cha `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
    document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `required`. Kumbuka, unahitaji tu kuongeza neno `required` ndani ya lebo ya kipengele cha `input`.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('required'));
```

Thamani haipaswi kutolewa kwa sifa ya kipengele `required`.

```js
assert.strictEqual(document.querySelector('input')?.getAttribute('required'), '');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
