---
id: 5ef9b03c81a63668521804d0
title: Hatua 27
challengeType: 0
dashedName: step-27
---

# --description--

Ili kuweka msisitizo kwenye neno au usemi fulani, unaweza kutumia kipengele cha `em`.

Sisistiza neno `love` katika kipengele cha `figcaption` kwa kukifunika na kipengele cha msisitizo cha `em`.

# --hints--

Kipengele chako cha msisitizo cha `em` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('em'));
```

Kipengele chako cha msisitizo cha `em` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/em\>/));
```

Umefuta kipengele cha `figcaption` au kinakosa lebo ya kufungua au ya kufunga.

```js
assert(document.querySelector('figcaption') && code.match(/<\/figcaption\>/));
```

Kipengele chako cha msisitizo cha `em` kinapaswa kuzunguka maandishi ya `love`. Ama umeacha maandishi hayo au kuna makosa ya tahajia.

```js
assert(
  document.querySelector('figcaption > em').innerText.toLowerCase() === 'love'
);
```

Maandishi ya `figcaption` yanapaswa kuwa `Cats love lasagna`. Kagua makosa ya tahajia na uhakikishe nafasi zinazohitajika zipo karibu na lebo za kufungua na kufunga za kipengele cha `em`.

```js
assert(
  document
    .querySelector('figcaption')
    .innerText.replace(/\s+/gi, ' ')
    .match(/cats love lasagna\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
--fcc-editable-region--
          <figcaption>Cats love lasagna.</figcaption>
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
