---
id: 5ef9b03c81a63668521804d4
title: Hatua 34
challengeType: 0
dashedName: step-34
---

# --description--

Kipengele cha `strong` kinatumika kuonyesha kwamba maandishi fulani ni ya umuhimu mkubwa au ya dharura.

Katika `figcaption` uliyoongeza, onyesha kwamba `hate` ni ya umuhimu mkubwa kwa kuizunguka katika kipengele cha `strong`.

# --hints--

Kipengele chako cha `strong` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('strong'));
```

Kipengele chako cha `strong` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/strong\>/));
```

Kipengele chako cha `strong` kinapaswa kuzunguka neno `hate` katika maandishi ya `Cats hate other cats.` Umeacha maandishi au kuna makosa ya tahajia.

```js
assert(
  document
    .querySelectorAll('figcaption')[1]
    .querySelector('strong')
    .innerText.trim().toLowerCase() === 'hate'
);
```

Maandishi ya `figcaption` yanapaswa kuwa `Cats hate other cats.` Kagua makosa ya tahajia na uhakikishe nafasi zinazohitajika zipo karibu na lebo za kufungua na kufunga za kipengele cha `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert(
  secondFigCaption &&
    secondFigCaption.innerText
      .replace(/\s+/gi, ' ')
      .trim()
      .match(/cats hate other cats\.?/i)
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
        
          <figcaption>Cats hate other cats.</figcaption>  
        
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
