---
id: 5ef9b03c81a63668521804d6
title: Hatua 37
challengeType: 0
dashedName: step-37
---

# --description--

Sasa utaongeza fomu ya mtandao kukusanya taarifa kutoka kwa watumizi.

Kipengele cha `form` hutumika kupata taarifa kutoka kwa mtumizi kama jina lao, barua pepe, na maelezo mengine.

Baada ya kichwa cha sehemu cha `Cat Form`, ongeza kipengele cha `form`.

# --hints--

Kipengele chako cha `form` kinapaswa kuwa na lebo ya kufungua na lebo ya kufunga. Huenda unakosa moja au zote mbili za lebo zinazohitajika, au unazoweka kwa mpangilio usio sahihi.

```js
assert(document.querySelector('form') && code.match(/<\/form>/g));
```

Lebos za kipengele chako cha `form` haziko kwa mpangilio sahihi.

```js
const noSpaces = code.replace(/\s/g, '');
assert(noSpaces.indexOf('<form>') < noSpaces.indexOf('</form>'));
```

Kipengele cha `form` kilichopangwa ndani ya kipengele cha `section` kinapaswa kuwa chini ya kipengele cha `h2`. Una kipengele cha `h2` na kipengele cha `form` kwa mpangilio usio sahihi.

```js
assert(document.querySelector('form').previousElementSibling.nodeName === 'H2');
```

Kipengele cha `form` hakipaswi kuwa na maudhui. Ondoa vipengele vyovyote vya HTML au maandishi kati ya lebo za kipengele cha `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
--fcc-editable-region--
        <h2>Cat Form</h2>
        
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
