---
id: 5ef9b03c81a63668521804e3
title: Hatua 60
challengeType: 0
dashedName: step-60
---

# --description--

Ongeza kisanduku cha kuchagua kingine baada ya ile uliyoiweka hivi karibuni. Thamani ya sifa ya kipengele `id` inapaswa kuwa `lazy` na thamani ya sifa ya kipengele `name` inapaswa kuwa ile ile kama kisanduku cha kuchagua cha mwisho.

Pia ongeza kipengele cha `label` upande wa kulia wa kisanduku kipya cha kuchagua chenye maandishi `Lazy`. Hakikisha unahusisha kipengele cha `label` na kisanduku kipya cha kuchagua kwa kutumia sifa ya kipengele `for`.

# --hints--

Unahitaji kuongeza kisanduku kipya cha kuchagua.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'), 2);
```

Kisanduku chako kipya cha kuchagua kinapaswa kuwa na sifa ya kipengele `id` yenye thamani `lazy` na sifa ya kipengele `name` yenye thamani `personality`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')];
assert(
  checkboxes.some(
    (checkbox) =>
      checkbox.id === 'lazy' && checkbox.getAttribute('name') === 'personality'
  )
);
```

Kisanduku chako kipya cha kuchagua kinapaswa kuwa baada ya cha kwanza. Umeviweka kwa mpangilio usio sahihi.

```js
const checkboxes = [...document.querySelectorAll('input[type="checkbox"]')].map(
  (checkbox) => checkbox.id
);
assert(checkboxes.indexOf('loving') < checkboxes.indexOf('lazy'));
```

Upande wa kulia wa kisanduku chako kipya cha kuchagua, kunapaswa kuwepo kipengele cha `label` chenye maandishi `Lazy`.

```js
const nextElementSibling = document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling;
assert.isOk(nextElementSibling?.nodeName === 'LABEL' && nextElementSibling?.innerText?.replace(/\s+/g, '')?.match(/^Lazy$/i)
);
```

`label` mpya inapaswa kuwa na sifa ya kipengele `for` yenye thamani ile ile kama sifa ya kipengele `id` ya kisanduku kipya cha kuchagua. Huenda umeacha thamani hiyo au umefanya makosa ya tahajia.

```js
assert.strictEqual(
  document.querySelectorAll('input[type="checkbox"]')?.[1]?.nextElementSibling?.getAttribute('for'),
    'lazy'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>
          <fieldset>
            <legend>What's your cat's personality?</legend>
--fcc-editable-region--
            <input id="loving" type="checkbox" name="personality"> <label for="loving">Loving</label>
--fcc-editable-region--
          </fieldset>
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
