---
id: 5ef9b03c81a63668521804d7
title: Hatua 38
challengeType: 0
dashedName: step-38
---

# --description--

Sifa ya kipengele `action` inaonyesha wapi data za fomu zinapaswa kutumwa.

Hapa kuna mfano wa kipengele cha `form` chenye sifa ya kipengele `action`:

```html
<form action="/submit-url"></form>
```

Katika mfano huu, `action="/submit-url"` inamwambia kivinjari kwamba data za fomu zinapaswa kutumwa kwenye njia ya `/submit-url`.

Ongeza sifa ya kipengele `action` yenye thamani `https://freecatphotoapp.com/submit-cat-photo` kwenye kipengele cha `form`.

# --hints--

Kipengele chako cha `form` kinapaswa kuwa na lebo ya kufungua na lebo ya kufunga kwa mpangilio sahihi. Huenda unakosa moja au zote mbili za lebo zinazohitajika, au unazo kwa mpangilio usio sahihi.

```js
const noSpaces = code.replace(/\s/g, '');
assert(
  document.querySelector('form') &&
    code.match(/<\/form>/g) &&
    noSpaces.indexOf('<form') < noSpaces.indexOf('</form>')
);
```

Kipengele chako cha `form` kilichopangwa ndani ya kipengele cha `section` kinapaswa kuwa chini ya kipengele cha `h2`. Una kipengele cha `h2` na kipengele cha `form` kwa mpangilio usio sahihi.

```js
assert.equal(document.querySelector('form')?.previousElementSibling?.nodeName, 'H2');
```

Kipengele chako cha `form` hakipaswi kuwa na maudhui yoyote. Ondoa vipengele vyovyote vya HTML au maandishi kati ya lebo za kipengele cha `form`.

```js
assert.lengthOf(document.querySelector('form')?.innerHTML?.trim(), 0);
```

Kipengele chako cha `form` hakina sifa ya kipengele `action`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za kipengele.

```js
const form = document.querySelector('form');
assert(form.hasAttribute('action'));
```

Kipengele chako cha `form` kinapaswa kuwa na sifa ya kipengele `action` yenye thamani `https://freecatphotoapp.com/submit-cat-photo`.

```js
const form = document.querySelector('form');
assert(
  form
    .getAttribute('action')
    .match(/^https:\/\/freecatphotoapp\.com\/submit-cat-photo$/)
);
```

Ingawa umeweka sifa ya kipengele `action` kwa URL sahihi, inashauriwa kila mara kuzunguka thamani ya sifa ya kipengele na alama za nukuu.

```js
assert(
  !/\<form\s+action\s*=\s*https:\/\/freecatphotoapp\.com\/submit-cat-photo/.test(
    code
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
--fcc-editable-region--
        <form>
        </form>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```
