---
id: 5efb2c990dc218d6c85f89b2
title: Hatua 45
challengeType: 0
dashedName: step-45
---

# --description--

Ingawa uliweka kitufe chako chini ya ingizo la maandishi, vinaonekana pembeni mwa kila mmoja kwenye ukurasa. Hii ni kwa sababu vipengele vya `input` na `button` ni <dfn>vipengele vya mstari mmoja</dfn>, ambavyo havionekani kwenye mistari mipya.

Kitufe ulichokiweka kita wasilisha fomu kwa chaguo-msingi. Hata hivyo, kutegemea tabia ya chaguo-msingi kunaweza kusababisha mkanganyiko. Ongeza sifa ya kipengele ya `type` yenye thamani `submit` kwenye `button` ili kufafanua kuwa ni kitufe cha kuwasilisha.

# --hints--

Kipengele chako cha `button` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelector('button'));
```

Kipengele chako cha `button` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/button\>/));
```

Kipengele chako cha `button` hakina sifa ya kipengele ya `type`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua.

```js
assert.isTrue(document.querySelector('button')?.hasAttribute('type'));
```

Kipengele chako cha `button` kinapaswa kuwa na sifa ya kipengele ya `type` yenye thamani `submit`. Huenda umeacha thamani au umefanya makosa ya tahajia. Kumbuka kuwa thamani za sifa za kipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert.match(
  document.querySelector('button')?.getAttribute('type'),
  /^submit$/i
);
```

Ingawa umeweka sifa ya kipengele ya `type` ya kipengele cha `button` kuwa `submit`, inashauriwa kila mara kuzunguka thamani ya sifa ya kipengele na alama za nukuu.

```js
assert(!/\<\s*button\s+type\s*=\s*submit/i.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
--fcc-editable-region--
          <button>Submit</button>
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
