---
id: 5f05a1d8e233dff4a68508d8
title: Hatua 49
challengeType: 0
dashedName: step-49
---

# --description--

Unda kitufe cha radio kingine chini ya kile cha kwanza. Kipange ndani ya kipengele cha `label` chenye `Outdoor` kama maandishi ya `label`. Mpe kitufe cha radio sifa ya `id` yenye thamani ya `outdoor`.

# --hints--

Kipengele cha `input` cha kitufe cha radio cha asili cha `Indoor` pekee ndicho kinapaswa kuwa na `id` iliyowekwa kuwa `indoor`. Unaweza kuanzisha upya hatua ili kupata HTML ya asili ikiwa inahitajika.

```js
assert(document.querySelectorAll('#indoor').length < 2);
```

Haupaswi kufanya mabadiliko yoyote kwa kitufe cha radio cha `Indoor`. Unaweza kuanzisha upya hatua ili kupata HTML ya asili ikiwa inahitajika.

```js
const indoorInput = document.querySelectorAll('#indoor');
assert(
  indoorInput.length == 1 &&
  indoorInput[0]?.nodeName?.toUpperCase() === 'INPUT' &&
  indoorInput[0]?.type === 'radio' &&
  code.match(/<label\s*>\s*<input [^>]*(id=["']?indoor["']?)[^>]*>\s*Indoor\s*<\/label>/)
);
```

Unapaswa kuongeza kipengele kipya cha `input` kilichopangwa ndani ya kipengele kipya cha `label`. Hakikisha `label` mpya ina lebo ya kufungua na lebo ya kufunga.

```js
assert(document.querySelectorAll('label input').length === 2);
```

Haupaswi kuongeza sifa yoyote kwa lebo ya kufungua ya `label`.

```js
assert(code.match(/<label\s*>/g).length === 2);
```

`input` mpya inapaswa kuwa na sifa ya `id` yenye thamani ya `outdoor`.

```js
assert(document.querySelector('label input[id="outdoor"]'));
```

`input` mpya inapaswa kuwa na sifa ya `type` yenye thamani ya `radio`.

```js
assert(document.querySelector('label input[id="outdoor"][type="radio"]'));
```

Haupaswi kuongeza vipengele vipya zaidi ya `input` kilichopangwa ndani ya `label`.

```js
assert(document.querySelector('label input[id="outdoor"]:only-child'));
```

Hakuna maandishi yanayopaswa kuwa kushoto ya `input` mpya.

```js
const outdoorBtn = document.querySelector('label input[id="outdoor"]');
assert(
    outdoorBtn?.previousSibling?.nodeName !== '#text' ||
    outdoorBtn?.previousSibling?.nodeValue?.replace(/\s/g, '') === ''
);
```

Maandishi ya lebo kwa kitufe chako kipya cha radio lazima yawe hasa `Outdoor`. Hii ni pamoja na herufi kubwa na ndogo.

```js
assert(document.querySelector('label > input[id="outdoor"]')?.nextSibling?.nodeValue?.replace(/^\s+|\s+$/g, '') === 'Outdoor');
```

Kitufe chako kipya cha radio na lebo inapaswa kuwa mara moja chini/baada ya kitufe cha radio cha `Indoor` na lebo yake. Hakuna lebo nyingine zinapaswa kuwa kati yao.

```js
assert(code.match(/<\/label>\s*<label\s*>\s*<input [^>]+>\s*Outdoor/i));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio"> Indoor</label>
          
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
