---
id: 5f07fb1579dc934717801375
title: Hatua 35
challengeType: 0
dashedName: step-35
---

# --description--

Ni wakati wa kuongeza sehemu mpya. Ongeza kipengele cha tatu cha `section` chini ya kipengele cha pili cha `section`.

# --hints--

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufungua. Lebo za kufungua zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelectorAll('section').length >= 3);
```

Unapaswa kuongeza lebo moja tu ya kufungua `section`. Tafadhali ondoa ziada zozote.

```js
assert(document.querySelectorAll('section').length === 3);
```

Kipengele chako cha `section` kinapaswa kuwa na lebo ya kufunga. Lebo za kufunga zina `/` mara tu baada ya herufi ya `<`.

```js
assert(code.match(/<\/section>/g).length >= 3);
```

Unapaswa kuongeza lebo moja tu ya kufunga `section`. Tafadhali ondoa ziada zozote.

```js
assert(code.match(/<\/section>/g).length === 3);
```

Vipengele vyote vya `section` vinapaswa kuwa kati ya lebo za kufungua na kufunga za kipengele cha `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const sectionElemsFound = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert(sectionElemsFound.length === 3);
```

Kipengele cha mwisho cha `section` hakipaswi kuwa na maudhui. Ondoa vipengele vyovyote vya HTML au maandishi ndani ya kipengele cha `section`.

```js
assert.lengthOf(document.querySelectorAll('main > section')?.[2]?.children, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
