---
id: 5f1a80975fc4bcae0edb3497
title: Hatua 51
challengeType: 0
dashedName: step-51
---

# --description--

Ikiwa utachagua kitufe cha radio cha `Indoor` na kuwasilisha fomu, data ya fomu kwa kitufe hicho itategemea sifa zake za `name` na `value`. Kwa kuwa vitufe vyako vya radio havina sifa ya `value`, data ya fomu itajumuisha `indoor-outdoor=on`, ambayo haifai wakati una vitufe vingi.

Ongeza sifa ya `value` kwa vitufe vyote vya radio. Kwa urahisi, weka sifa ya `value` ya kitufe kuwa na thamani ile ile kama sifa yake ya `id`.

# --hints--

Vitufe vyote vya radio vinapaswa bado kuwa kati ya lebo za kufungua na kufunga za kipengele cha `label`.

```js
const labelChildNodes = [...document.querySelectorAll('form > label')].map(
  (node) => node.childNodes
);
assert(
  labelChildNodes.filter((childNode) => childNode[0].nodeName === 'INPUT')
    .length === 2
);
```

Vitufe vyote vya radio vinapaswa kuwa na sifa ya `value`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
const radioButtons = [...document.querySelectorAll('input[type="radio"]')];
assert(radioButtons.every((btn) => btn.hasAttribute('value')));
```

Vitufe vyote vya radio vinapaswa kuwa na sifa ya `type`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa.

```js
const radioButtons = [...document.querySelectorAll('input')];
assert(radioButtons.every((btn) => btn.hasAttribute('type')));
```

Sifa ya `Indoor` ya kitufe cha radio cha `value` inapaswa kuwekwa kuwa `indoor`. Umeacha thamani au kuna makosa ya tahajia. Kumbuka kuwa thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
const indoorRadioButton = document.querySelector('#indoor');
assert(indoorRadioButton.getAttribute('value').match(/^indoor$/));
```

Sifa ya `Outdoor` ya kitufe cha radio cha `value` inapaswa kuwekwa kuwa `outdoor`. Umeacha thamani au kuna makosa ya tahajia. Kumbuka kuwa thamani za sifa zinapaswa kuzungukwa na alama za nukuu.

```js
const outdoorRadioButton = document.querySelector('#outdoor');
assert(outdoorRadioButton.getAttribute('value').match(/^outdoor$/));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <label><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
          <label><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label>
--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```
