---
id: 671141feba228a35cefba82d
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Badilisha maandishi yaliyopo `cute cats` kuwa kipengele cha nanga kinachounganisha na:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`

# --hints--

Kipengele kipya cha nanga kinapaswa kuwepo katika kipengele cha kwanza cha `p`.

```js
assert.exists(document.querySelector('p')?.querySelector('a'));
```

Maandishi ya nanga yanapaswa kuwa `cute cats`.

```js
assert.strictEqual(document.querySelector('a')?.innerText.trim(), "cute cats");
```

Sifa ya kipengele `href` ya kipengele kipya cha nanga inapaswa kuwa `https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg`.

```js
assert.strictEqual(document.querySelector('a')?.href, "https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg");
```

Inapaswa kuwepo na mfano mmoja tu wa usemi `cute cats` katika msimbo wako.

```js
assert.strictEqual(code.indexOf("cute cats"),code.lastIndexOf("cute cats")); 
```

Maandishi ya kipengele cha `p` yanapaswa kubaki kuwa `Everyone loves cute cats online!`.

```js
assert.strictEqual(document.querySelector('p')?.innerText.trim(), "Everyone loves cute cats online!");
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
