---
id: 671141d8e32fe934c26fa1be
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Ongeza maneno `See more ` kabla ya kipengele cha nanga na ` in our gallery` baada ya kipengele cha nanga.

# --hints--

Unapaswa kutumia kipengele cha nanga kilichopo, usitengeneze kipengele kipya.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

Maandishi ndani ya kipengele cha nanga hayapaswi kubadilika, lazima yawe `cat photos`.

```js
assert.strictEqual(document.querySelector('a').innerText, "cat photos");
```

Unapaswa kuwa na `See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery` katika msimbo wako.

```js
assert.match(code, /See more <a href="https:\/\/freecatphotoapp\.com">cat photos<\/a> in our gallery/)
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com">cat photos</a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
