---
id: 7cf9b03d81a65668421804c3
title: Hatua 41
challengeType: 0
dashedName: step-41
---

# --description--

Ili data ya fomu ipatikane na mahali panapobainishwa katika sifa ya kipengele `action`, lazima uweke sifa ya kipengele `name` kwenye uwanja wa maandishi na uweke thamani inayowakilisha data inayowasilishwa.

Hapa kuna mfano wa kipengele cha `input` chenye sifa ya kipengele `name`:

```html
<input type="text" name="name">
```

Ongeza sifa ya kipengele `name` yenye thamani `catphotourl` kwenye uwanja wako wa maandishi.

# --hints--

Umefuta kipengele chako cha `input` au kina sintaksia isiyo sahihi. Thamani zote za sifa za vipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert(document.querySelectorAll('input').length);
```

`form` yako inapaswa kuwa na kipengele cha `input` pekee. Futa vipengele vingine vya HTML au maandishi yaliyomo ndani ya kipengele cha `form`.

```js
assert.isTrue(
  document.querySelector('form')?.children?.length === 1 &&
  document.querySelector('form')?.innerText?.trim()?.length === 0
);
```

Kipengele chako cha `input` hakina sifa ya kipengele `name`. Hakikisha kuna nafasi baada ya jina la lebo ya kufungua na/au kuna nafasi kabla ya majina yote ya sifa za vipengele.

```js
assert.isTrue(document.querySelector('input')?.hasAttribute('name'));
```

Kipengele chako cha `input` kinapaswa kuwa na sifa ya kipengele `name` yenye thamani `catphotourl`. Umeacha thamani au kuna makosa ya tahajia. Kumbuka thamani za sifa za vipengele zinapaswa kuzungukwa na alama za nukuu.

```js
assert.match(
  document.querySelector('input')?.getAttribute('name'),
  /^catphotourl$/i
);
```

Ingawa umeweka sifa ya kipengele `input` ya kipengele cha `name` kuwa `catphotourl`, inashauriwa kila mara kuzunguka thamani ya sifa ya kipengele na alama za nukuu.

```js
assert(!/\<\s*input\s+.*\s*=\s*catphotourl/.test(code));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <input type="text">
--fcc-editable-region--
        </form>
      </section>
    </main>
  </body>
</html>
```
