---
id: 68e4276483b0c71658aab1d1
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Wenn Sie einen neuen String erstellen möchten, bei dem alle Buchstaben in Kleinbuchstaben sind, können Sie die `toLowerCase()`-Methode verwenden.

Beginnen Sie damit, eine Variable namens `lowerCaseInput` zu erstellen und weisen Sie ihr das Ergebnis der Verwendung der `toLowerCase()`-Methode auf der `cleanedInput`-Variable zu.

Geben Sie dann den String `"Result of using the toLowerCase() method:"` in der Konsole aus.

Fügen Sie schließlich eine weitere `console.log()` für die `lowerCaseInput`-Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine `lowerCaseInput`-Variable haben.

```js
assert.exists(lowerCaseInput);
```

Ihre `lowerCaseInput`-Variable sollte ein String sein.

```js
assert.isString(lowerCaseInput);
```

Sie sollten `cleanedInput.toLowerCase()` Ihrer `lowerCaseInput`-Variable zuweisen.

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

Sie sollten den String `"Result of using the toLowerCase() method:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

Sie sollten die `lowerCaseInput`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
