---
id: 68e4276483b0c71658aab1d0
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Es wird Situationen geben, in denen Sie einen neuen String mit einer anderen Groß- und Kleinschreibung erstellen müssen.

Um einen neuen String zu erstellen, bei dem alle Buchstaben in Großbuchstaben umgewandelt sind, können Sie die `toUpperCase()`-Methode verwenden.

Hier ist ein Beispiel:

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

Erstellen Sie eine Variable namens `upperCaseInput` und weisen Sie ihr das Ergebnis der Verwendung der `toUpperCase()`-Methode auf die `cleanedInput`-Variable zu.

# --hints--

Sie sollten eine `upperCaseInput`-Variable haben.

```js
assert.exists(upperCaseInput);
```

Ihre `upperCaseInput`-Variable sollte ein String sein.

```js
assert.isString(upperCaseInput);
```

Sie sollten `cleanedInput.toUpperCase()` Ihrer `upperCaseInput`-Variable zuweisen.

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
