---
id: 68e4276483b0c71658aab1d2
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Es wird Zeiten geben, in denen Sie mehrere Methoden verwenden müssen, um einen neuen String zu erstellen. In den nächsten Schritten werden Sie einige der Methoden anwenden, die Sie bisher gelernt haben, um einen neuen String zu erstellen, der die CamelCase-Benennungskonvention verwendet.

Beginnen Sie damit, eine Variable namens `lowercaseWord` zu erstellen und ihr den String `"camelcase"` zuzuweisen.

Unterhalb dieser `lowercaseWord`-Variable erstellen Sie eine weitere Variable namens `camelCasedVersion` und weisen ihr vorerst einen leeren String zu.

Unterhalb dieser `camelCasedVersion`-Variable fügen Sie eine `console.log()` mit dem String `"Camel cased version:"` hinzu, gefolgt von einer weiteren `console.log()` für die `camelCasedVersion`-Variable.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten eine Variable namens `lowercaseWord` haben.

```js
assert.exists(lowercaseWord);
```

Ihr `lowercaseWord` sollte ein String sein.

```js
assert.isString(lowercaseWord);
```

Sie sollten den String `"camelcase"` Ihrer `lowercaseWord`-Variable zuweisen.

```js
assert.equal(lowercaseWord, "camelcase");
```

Sie sollten eine `camelCasedVersion`-Variable haben.

```js
assert.exists(camelCasedVersion);
```

Ihre `camelCasedVersion`-Variable sollte ein String sein.

```js
assert.isString(camelCasedVersion);
```

Sie sollten Ihrer `camelCasedVersion`-Variable einen leeren String zuweisen.

```js
assert.isEmpty(camelCasedVersion);
```

Sie sollten den String `"Camel cased version:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

Sie sollten die `camelCasedVersion`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
