---
id: 68e4276483b0c71658aab1cc
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Ihr `userInput`-Wert enthält an beiden Enden des Strings einige zusätzliche Leerzeichen. Um einen neuen String zu erstellen, bei dem diese Leerzeichen entfernt sind, können Sie die `trim()`-Methode verwenden.

Hier ist ein Beispiel:

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

Erstellen Sie eine Variable mit dem Namen `cleanedInput` und weisen Sie ihr das Ergebnis der Anwendung der `trim()`-Methode auf die `userInput`-Variable zu.

# --hints--

Sie sollten eine `cleanedInput`-Variable haben.

```js
assert.exists(cleanedInput);
```

Ihre `cleanedInput`-Variable sollte ein String sein.

```js
assert.isString(cleanedInput);
```

Sie sollten `userInput.trim()` Ihrer `cleanedInput`-Variable zuweisen.

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
