---
id: 68e4276483b0c71658aab1ce
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Es wird Situationen geben, in denen Sie nur Leerzeichen am Anfang des Strings entfernen möchten. Hier ist die `trimStart()`-Methode sehr nützlich.

Hier ist ein Beispiel:

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

Denken Sie daran, dass Strings unveränderlich sind. Diese Methode erstellt einen neuen String mit dem Ergebnis. Der ursprüngliche wird nicht verändert.

Erstellen Sie eine Variable `trimmedStart` und weisen Sie ihr das Ergebnis der Anwendung der `trimStart()`-Methode auf die Variable `userInput` zu.

# --hints--

Sie sollten eine `trimmedStart`-Variable haben.

```js
assert.exists(trimmedStart);
```

Ihre `trimmedStart`-Variable sollte ein String sein.

```js
assert.isString(trimmedStart);
```

Sie sollten `userInput.trimStart()` Ihrer `trimmedStart`-Variable zuweisen.

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
