---
id: 68e4276483b0c71658aab1ca
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit verschiedenen String-Methoden, die zum Formatieren von Strings verwendet werden.

Beginnen Sie damit, eine Variable namens `userInput` zu erstellen und ihr den folgenden Wert zuzuweisen: `"   Hello World!   "`.

# --hints--

Sie sollten eine `userInput`-Variable haben.

```js
assert.exists(userInput);
```

Ihre `userInput`-Variable sollte ein String sein.

```js
assert.isString(userInput);
```

Sie sollten `"   Hello World!   "` Ihrer `userInput`-Variable zuweisen. Achten Sie darauf, dass vor `H` drei Leerzeichen und nach `!` drei Leerzeichen stehen.

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
