---
id: 68e4276483b0c71658aab1cd
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Jetzt ist es Zeit, das Ergebnis zu sehen.

Beginnen Sie damit, einen `console.log()` mit dem String `"Result of trimming whitespace from both ends:"` hinzuzufügen.

Fügen Sie dann darunter `console.log()` einen weiteren `console.log()` mit der `cleanedInput` Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"Result of trimming whitespace from both ends:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

Sie sollten die `cleanedInput` Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
