---
id: 68e4276483b0c71658aab1d3
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Wie Sie sich aus früheren Lektionen und Workshops erinnern, können Sie Teile eines Strings mit der `slice()`-Methode extrahieren.

Hier eine Erinnerung:

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

Sie sollten die bestehende `camelCasedVersion`-Zuweisung ändern, indem Sie den leeren String durch `lowercaseWord.slice(0, 5)` ersetzen.

Jetzt sollten Sie das Wort `"camel"` in der Konsole sehen.

# --hints--

Ihre `camelCasedVersion`-Variable sollte kein leerer String sein.

```js
assert.isNotEmpty(camelCasedVersion);
```

Sie sollten `lowercaseWord.slice(0, 5)` der `camelCasedVersion`-Variable zuweisen.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

Sie sollten NICHT den wörtlichen String `"camel"` der `camelCasedVersion` zuweisen. Bitte lesen Sie die Anweisungen noch einmal für den korrekten Ansatz.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
