---
id: 68ed6079fe0e3e98b94a1b75
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Fügen Sie als Nächstes einen `console.log()` mit dem String `"Original input:"` hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"Original input:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
