---
id: 68f1179eaea3bb59ba434386
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Um Leerzeichen am Ende eines Strings zu entfernen, können Sie die `trimEnd()`-Methode verwenden.

Hier ist ein Beispiel:

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

Erstellen Sie eine Variable namens `trimmedEnd` und weisen Sie ihr das Ergebnis der Verwendung der `trimEnd()`-Methode auf der `userInput`-Variable zu.

# --hints--

Sie sollten eine `trimmedEnd`-Variable haben.

```js
assert.exists(trimmedEnd);
```

Ihre `trimmedEnd`-Variable sollte ein String sein.

```js
assert.isString(trimmedEnd);
```

Sie sollten `userInput.trimEnd()` Ihrer `trimmedEnd`-Variable zuweisen.

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
