---
id: 68f3eeaa65614107ecd5f1f2
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Fügen Sie unter Ihrem ersten `console.log()` eine weitere `console.log()` hinzu, um die `userInput`-Variable in der Konsole auszugeben.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten die `userInput`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
