---
id: 68f3f5ca1d2a490c97f14f03
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Jetzt ist es Zeit, die Ergebnisse des Abschneidens von Leerzeichen am Ende des Strings zu sehen.

Fügen Sie unter Ihrer `trimmedEnd`-Variable eine `console.log()` mit dem String `"After using the trimEnd() method, trailing spaces removed:"` hinzu.

Dann fügen Sie unter dieser `console.log()` eine weitere `console.log()` mit der `trimmedEnd`-Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"After using the trimEnd() method, trailing spaces removed:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

Sie sollten die `trimmedEnd`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
