---
id: 68f3f890010c870e17bbe754
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Jetzt ist es Zeit, die Ergebnisse zu sehen.

Fügen Sie unter Ihrer `upperCaseInput`-Variable eine `console.log()` mit dem String `"Result of using the toUpperCase() method:"` hinzu.

Dann fügen Sie unter dieser `console.log()` eine weitere `console.log()` mit der `upperCaseInput`-Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"Result of using the toUpperCase() method:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

Sie sollten die `upperCaseInput`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
