---
id: 68f3f3d65d51d80ac0b412fc
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, die Ergebnisse des Abschneidens von Leerzeichen am Anfang des Strings zu sehen.

Fügen Sie unter Ihrer `trimmedStart`-Variable eine `console.log()` mit dem String `"After using the trimStart() method, leading spaces removed:"` hinzu.

Dann fügen Sie unter dieser `console.log()` eine weitere `console.log()` mit der `trimmedStart`-Variable hinzu.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Sie sollten den String `"After using the trimStart() method, leading spaces removed:"` in der Konsole ausgeben.

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

Sie sollten die `trimmedStart`-Variable in der Konsole ausgeben.

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
