---
id: 68f405c8c9df44132c8fe1bb
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Wie Sie sich aus früheren Lektionen erinnern, beginnt die CamelCase-Konvention mit einem Kleinbuchstaben für das erste Wort, und jedes folgende Wort beginnt mit einem Großbuchstaben, gefolgt von Kleinbuchstaben.

Hier sind einige Erinnerungsbeispiele:

```md
freeCodeCamp
camelCase
learningIsFun
```

Das zweite Wort in der `lowercaseWord`-Variablen ist `"case"`. Um auf den `c` in diesem Wort zuzugreifen, können Sie `lowercaseWord[5]` verwenden.

Aktualisieren Sie die `camelCasedVersion`-Zuweisung, um den `+`-Operator zu verwenden, der `lowercaseWord.slice(0, 5)` mit dem Ergebnis der Verwendung der richtigen Methode zur Umwandlung von Strings in Großbuchstaben auf `lowercaseWord[5]` verknüpft.

Jetzt sollten Sie `camelC` in der Konsole sehen.

# --hints--

Sie sollten `lowercaseWord.slice(0, 5)` mit `lowercaseWord[5].toUpperCase()` verknüpfen.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

Sie sollten NICHT den Literalstring `"camelC"` der `camelCasedVersion` zuweisen. Bitte lesen Sie die Anweisungen noch einmal für den korrekten Ansatz.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

