---
id: 68f40925d810ae14b2d5a169
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Im Moment enthält Ihre Variable `camelCasedVersion` den Wert von `camelC`. Das gewünschte Ergebnis ist jedoch `camelCase`. 

Um die verbleibenden `ase` Buchstaben zu erhalten, können Sie die `slice()`-Methode mit einem negativen Indexwert verwenden. Denken Sie daran, dass die Verwendung negativer Indexwerte der `slice()`-Methode sagt, vom Ende aus zu extrahieren. 

Hier ist ein Beispiel zur Erinnerung:

```js
const str = "abcd";
str.slice(-2); // cd
```

Aktualisieren Sie die `camelCasedVersion`-Zuweisung, sodass sie auch `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` mit dem Ergebnis der Verwendung der `slice()`-Methode auf `lowercaseWord` verknüpft. Übergeben Sie für die `-3`-Methode die Zahl `slice()`.

Und mit dieser letzten Änderung ist Ihr Workshop abgeschlossen!

# --hints--

Sie sollten `lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` mit `lowercaseWord.slice(-3)` verknüpfen.

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

Sie sollten NICHT den Literal-String `"camelCase"` der `camelCasedVersion` zuweisen. Bitte lesen Sie die Anweisungen noch einmal durch, um den korrekten Ansatz zu verstehen.

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
