---
id: 660400f02cf3a76795b82d4e
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Першу річ яку треба налаштувати, це метод, який обчислює переміщення снаряда, яке є горизонтальним простором, пройденим від кидка до моменту, коли снаряд торкається землі.

Створіть метод `__calculate_displacement`, який має лише `self` як параметр, та поверніть переміщення снаряда.

Use the following formula to compute the projectile displacement: \\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \\]

Де $d$ — переміщення, $v$ — початкова швидкість, $\theta$ — кут, а $h$ — початкова висота снаряда. Для $g$ ви можете використовувати змінну `GRAVITATIONAL_ACCELERATION`.

Використайте методи `math.cos()` та `math.sin()` для тригонометричних функцій і `math.sqrt()` для обчислення квадратного кореня. Також ви повинні знати, що $x^y$ записується як `x ** y` у python. Також $\sin^2\theta$ означає, що значення, вислідне з синуса, потім зводиться у квадрат.

Пам’ятайте, що зі спотворенням імен потрібно викликати метод як `_Projectile__calculate_displacement`, якщо ви хочете протестувати, або використовувати його поза класом:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

You should declare a method called `__calculate_displacement` with `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

Метод `__calculate_displacement` повинен мати лише аргумент `self`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

Метод `__calculate_displacement` має повернути правильне значення.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)

--fcc-editable-region--

--fcc-editable-region--
```
