---
id: 660fcf3e1b9bb056b2edb567
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

На цьому етапі ви готові створити рядкове представлення. Почніть зі створення екземпляра класу `Projectile`. Під визначенням класу створіть змінну `ball` та призначте до неї виклик до `Projectile`, використовуючи `10, 3, 45` як аргументи.

Визначте метод `__str__` у класі, щоб надати вашому екземпляру належне рядкове представлення та змусити його повертати рядок у наступному форматі:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Він повинен починатися та закінчуватися символом нового рядка, кут повинен бути записаний як ціле число в градусах, а переміщення має бути надруковано з одним десятковим знаком.

Ви знайдете корисним `math.degrees` для перетворення кута з радіанів на градуси.

Коли ви будете готові, ви можете `print(ball)`, щоб перевірити свою функцію.

# --hints--

Ви повинні мати змінну `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

Змінна `ball` повинна мати значення `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

Представлення рядка для `Projectile(45, 45, 45)` має бути правильним.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

Рядкове представлення також має бути правильним для інших екземплярів.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Ви повинні мати виклик `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)

        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION

--fcc-editable-region--

--fcc-editable-region--

```
