---
id: 66335223f3e2a31d62d84367
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Повертаючись до калькулятора траєкторії снаряда, тепер ви створите новий клас, який приймає список координат та створює креслення траєкторії.

Створіть новий клас `Graph`, єкземпляр якого має бути створений за допомогою приватного атрибута `__coordinates`, де зберігається список координат. Оголосіть змінну класу `__slots__` і визначте спеціальні методи `__init__` та `__repr__`.

Також створіть екземпляр класу `Graph`, передавши змінну `coordinates` та призначивши її до нової змінної `graph`.

# --hints--

Створіть новий клас `Graph`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).has_class('Graph')`
    ))
})
```

Ви повинні мати змінну `__slots__` в межах класу.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_variable('__slots__')`
    ))
})
```

Змінна `__slots__` повинна мати значення `('__coordinates')` або `['__coordinates']`.

```js
({
    test: () => (runPython(
        `
        slots = _Node(_code).find_class('Graph').find_variable('__slots__')

        assert any(slots.is_equivalent(str) for str in ['__slots__ = ("__coordinates")', '__slots__ = ["__coordinates"]'])`
    ))
})
```

Ви повинні мати метод `__init__`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__init__')`
    ))
})
```

Метод `__init__` має створити екземпляр приватної властивості `__coordinates` із переданим значенням.

```js
({
    test: () => assert(runPython(
        `
        l = [(3, 4), (4, 1), (3, 9)]
        g = Graph(l)
        g._Graph__coordinates == l
        `
    ))
})
```

Ви повинні мати метод `__repr__`.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_class('Graph').has_function('__repr__')`
    ))
})
```

Метод `__repr__` має повернути правильний рядок.

```js
({
    test: () => assert(runPython(
        `
        ball = Projectile(10, 3, 45)
        coords = ball.calculate_all_coordinates()
        repr(Graph(coords)) == f'Graph({coords})'
        `
    ))
})
```

Створіть екземпляр класу `Graph`, передавши змінну `coordinates` та призначивши її до змінної `graph`.

```js
({
    test: () => {
        runPython(`

        assert _Node(_code).has_variable('graph'), "variable graph is missing"

        assert _Node(_code).find_variable('graph').is_equivalent('graph = Graph(coordinates)'), "assignment to graph is wrong"
        `)
    }
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)

    def __str__(self):
        return f'''
Projectile details:
speed: {self.speed} m/s
height: {self.height} m
angle: {self.angle}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)

        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION

    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate

    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

    @property
    def height(self):
        return self.__height

    @property
    def angle(self):
        return round(math.degrees(self.__angle))

    @property
    def speed(self):
        return self.__speed

    @height.setter
    def height(self, n):
        self.__height = n

    @angle.setter
    def angle(self, n):
        self.__angle = math.radians(n)

    @speed.setter
    def speed(self, s):
       self.__speed = s

    def __repr__(self):
        return f"{self.__class__.__name__}({self.speed}, {self.height}, {self.angle})"

--fcc-editable-region--


ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()


--fcc-editable-region-- 
```
