---
id: 66a37f37ef5823a313de8c26
title: Crie uma órbita lunar
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

Neste laboratório, você criará uma animação simples da órbita da Lua ao redor da Terra usando HTML e CSS. A Terra estará no centro do sistema, enquanto a Lua girará ao redor dela.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve alinhar todos os elementos ao centro da página definindo a altura como `100%` da viewport e configurando um layout do flexbox no elemento `body`.
   
2. Você deve ter uma `div` com a classe `space`.

3. Dentro do elemento `.space`, devem haver mais dois elementos `div` com as classes `earth` e `orbit`, respectivamente, na ordem.

4. Dentro do elemento `.orbit`, deve haver outro `div` com a classe `moon`.
   
5. O elemento `div` com uma `class` de `space` deve ser centralizado na página e ter uma largura e altura de `200px`.

6. O elemento `div` com uma `class` de `space` deve usar posicionamento `relative`.

7. O elemento `.earth` deve usar posicionamento `absolute`. Posicione o centro do elemento `.earth` no ponto médio de seu elemento pai tanto no eixo vertical (top) quanto no eixo horizontal (left). Após isso, desloque o elemento `.earth` para trás pela metade de sua própria largura e altura, para centralizá-lo dentro de seu elemento pai, o elemento `.space`.

8. O elemento `.earth` deve ter largura e altura de `100px`.

9.  O elemento `.orbit` deve ter uma largura e altura de `200px`.

10.  O elemento `.orbit` deve ser posicionado usando posicionamento `absolute`. Seu canto inferior direito deve estar no centro do elemento `.space` usando uma propriedade `transform` que o desloca em `-50%` tanto no eixo vertical quanto no horizontal.

11. O caminho da órbita para a lua ao redor da Terra deve ser um círculo.

12. O elemento `.moon` deve ser posicionado usando posicionamento `absolute` e ter largura e altura de `30px`. O elemento `.moon` deve se posicionar no topo do elemento `.orbit` e estar centralizado horizontalmente.

13. Você deve ajustar ainda mais o posicionamento horizontal do elemento `.moon` movendo o elemento para a esquerda pela metade de sua largura.

14. O seletor `.earth` deve ter uma cor de fundo e um `border-radius` de `50%`.

15. O seletor `.moon` deve ter uma cor de fundo e um `border-radius` de `50%`.

16. Você deve definir uma animação `@keyframes orbit` que gira o elemento `.orbit` 360 graus ao redor de seu centro. Você deve aplicar essa animação ao elemento `.orbit` com duração de `5` segundos, uma função de tempo linear e iterações infinitas.

**Nota:** Certifique-se de vincular sua folha de estilos no seu HTML e aplicar seu CSS.

# --hints--

O conteúdo do elemento `body` deve estar centralizado na página definindo a altura como `100%` da altura da viewport, `display` como `flex`, `justify-content` e `align-items` como `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Você deve ter uma `div` principal com a classe `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

O elemento `.space` deve ter largura e altura de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

O elemento `.space` deve ter uma propriedade `position` com valor `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Você deve ter um `div` com uma classe `earth` dentro do `div` com uma `class` `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

O elemento `.earth` deve ser o primeiro filho do elemento `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

O elemento `.earth` deve ter largura e altura de `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

O elemento `.earth` deve usar posicionamento `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

O elemento `.earth` deve ter uma posição `top` e `left` de `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

O elemento `.earth` deve ser movido para trás e para cima em `50%` definindo a propriedade `transform` para `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Você deve ter um `div` com a classe `orbit` dentro do elemento `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

O elemento `.orbit` deve vir após o elemento `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

O elemento `.orbit` deve ter uma largura e altura de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

O elemento `.orbit` deve usar posicionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

O elemento `.orbit` deve ser posicionado no centro do elemento `.space` usando uma propriedade `transform` que o desloca em `-50%` tanto no eixo vertical quanto no horizontal.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Você deve ter um `div` com a classe `moon` dentro do elemento `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

O elemento `.moon` deve ter largura e altura de `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

O elemento `.moon` deve ser posicionado usando posicionamento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

A borda superior do elemento `.moon` deve ser alinhada com a borda superior de seu bloco de contenção.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Você deve mover a borda esquerda do elemento `.moon` para o centro do bloco de contenção (`50%` do lado esquerdo).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

O elemento `.moon` deve ser ajustado horizontalmente traduzindo-o para `-50%` ao longo do eixo x.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

O seletor `.earth` deve ter uma cor de fundo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

O seletor `.earth` deve ter `border-radius` de `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

O seletor `.moon` deve ter uma cor de fundo.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

O seletor `.moon` deve ter `border-radius` de `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Você deve criar uma tag `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

A nova regra `@keyframes` deve ser nomeada `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

A regra `@keyframes orbit` deve ter um seletor `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

A regra `@keyframes orbit` deve ter um seletor `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

O seletor `0%` deve ter uma propriedade `transform` definida como `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

O elemento `100%` deve vir depois do elemento `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

O seletor `100%` deve ter uma propriedade `transform` definida como `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

O seletor `.orbit` deve ter uma propriedade `animation` definida como `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

