---
id: 615f3cafd794015aa9547a6d
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

تذكّر أن استخدام الوسوم `h1` و `h2` وما يشابهها يحدد الهيكل الدلالي لصفحة HTML الخاصة بك. مع ذلك، يمكنك تعديل CSS لهذه العناصر للتحكم في التدفق البصري والترتيب الهرمي.

أنشئ قاعدة `h1` واضبط خاصية `font-weight` على القيمة `800`. هذا سيجعل نص `h1` أكثر عريضًا.

# --hints--

يجب أن تنشئ محدِّد `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

يجب أن يحتوي محدِّد `h1` على خاصية `font-weight` مضبوطة على القيمة `800`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.fontWeight === '800');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

--fcc-editable-region--

--fcc-editable-region--
```
