---
id: 615f2abbe7d18d49a1e0e1c8
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

تم تزويدك بقالب جاهز أساسي لـ HTML.

أنشئ عنصر `h1` داخل عنصر `body` وأضف له النص `Nutrition Facts`.

# --hints--

يجب إضافة عنصر `h1` جديد.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يكون عنصر `h1` داخل عنصر `body`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.localName, 'body');
```

يجب أن يحتوي عنصر `h1` على النص `Nutrition Facts`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Nutrition Facts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```

```css

```
