---
id: 615f3ed16592445e57941aec
title: الخطوة 16
challengeType: 0
dashedName: step-16
---

# --description--

تساعد الخطوط في فصل وتجميع المحتوى المهم، خاصةً عندما تكون المساحة محدودة.

أنشئ عنصر `div` أسفل عنصر `h1`، وأعطه خاصية `class` بقيمة `divider`.

# --hints--

يجب أن تنشئ عنصر `div` جديد.

```js
assert(document.querySelectorAll('div')?.length === 2);
```

يجب أن يحتوي عنصر `div` الجديد على خاصية `class` بقيمة `divider`.

```js
assert(document.querySelectorAll('div')?.[1]?.classList?.contains('divider'));
```

يجب أن يكون عنصر `.divider` داخل عنصر `.label`.

```js
assert(document.querySelector('.label')?.querySelector('.divider'));
```

يجب أن يأتي عنصر `.divider` بعد عنصر `h1`.

```js
assert(document.querySelector('.divider')?.previousElementSibling?.localName === 'h1');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <div class="label">
    <h1>Nutrition Facts</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
    <p>8 servings per container</p>
    <p>Serving size 2/3 cup (55g)</p>
  </div>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}

.label {
  border: 2px solid black;
  width: 270px;
  margin: 20px auto;
  padding: 0 7px;
}

h1 {
  font-weight: 800;
  text-align: center;
  margin: -4px 0;
}

p {
  margin: 0;
}
```
