---
id: 615f34948891834dd77655a6
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أضف عنصر `p` ثانيًا يحتوي على النص `Serving size 2/3 cup (55g)`.

# --hints--

يجب أن يكون لديك عنصر `p` ثاني.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

يجب أن يكون عنصر `p` الثاني داخل عنصر `body` الخاص بك.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

يجب أن يأتي عنصر `p` الثاني بعد عنصر `p` الموجود لديك.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

يجب أن يحتوي عنصر `p` الثاني على النص `Serving size 2/3 cup (55g)`.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```

```css

```
