---
id: acda2fb1324d9b0fa741e6b5
title: implementare uno strumento per confermare la fine di una stringa
challengeType: 26
dashedName: implement-a-string-ending-checker-function
---

# --description--

In questo laboratorio implementerai una funzione che verifica se una stringa termina con una data stringa obiettivo.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `confirmEnding` che prende due parametri: la stringa da controllare e la stringa con cui confrontarla.
2. La funzione dovrebbe restituire `true` se la prima stringa termina con la seconda stringa, e `false` altrimenti.
3. Non devi usare il metodo `.endsWith()`; invece, usa uno dei metodi substring di JavaScript per ottenere questo risultato.

# --hints--

Devi creare una funzione chiamata `confirmEnding`.

```js
assert.isFunction(confirmEnding);
```

`confirmEnding` deve prendere 2 parametri.

```js
assert.lengthOf(confirmEnding, 2);
```

`confirmEnding("Bastian", "n")` deve restituire `true`.

```js
assert.isTrue(confirmEnding('Bastian', 'n'));
```

`confirmEnding("Congratulation", "on")` deve restituire `true`.

```js
assert.isTrue(confirmEnding('Congratulation', 'on'));
```

`confirmEnding("Connor", "n")` deve restituire `false`.

```js
assert.isFalse(confirmEnding('Connor', 'n'));
```

`confirmEnding("Walking on water and developing software from a specification are easy if both are frozen", "specification")` deve restituire `false`.

```js
assert.isFalse(
  confirmEnding(
    'Walking on water and developing software from a specification are easy if both are frozen',
    'specification'
  )
);
```

`confirmEnding("He has to give me a new name", "name")` deve restituire `true`.

```js
assert.isTrue(confirmEnding('He has to give me a new name', 'name'));
```

`confirmEnding("Open sesame", "same")` deve restituire `true`.

```js
assert.isTrue(confirmEnding('Open sesame', 'same'));
```

`confirmEnding("Open sesame", "sage")` deve restituire `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'sage'));
```

`confirmEnding("Open sesame", "game")` deve restituire `false`.

```js
assert.isFalse(confirmEnding('Open sesame', 'game'));
```

`confirmEnding("If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing", "mountain")` deve restituire `false`.

```js
assert.isFalse(
  confirmEnding(
    "If you want to save our world, you must hurry. We don't know how much longer we can withstand the nothing",
    'mountain'
  )
);
```

`confirmEnding("Abstraction", "action")` deve restituire `true`.

```js
assert.isTrue(confirmEnding('Abstraction', 'action'));
```

Il tuo codice non deve usare il metodo integrato `.endsWith()` per risolvere il laboratorio.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.endsWith\(.*?\)\s*?;?/);
assert.notMatch(__helpers.removeJSComments(code), /\['endsWith'\]/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function confirmEnding(str, target) {
  return str.substring(str.length - target.length) === target;
}

confirmEnding('Bastian', 'n');
```
