---
id: 686ccc2c8b967e17ab18d593
title: implementare il metodo della bisezione
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

Il metodo della bisezione, noto anche come metodo di ricerca binaria, usa una ricerca binaria per trovare le radici di una funzione a valori reali. Funziona restringendo un intervallo in cui si trova la radice quadrata fino a convergere a un valore entro una tolleranza specificata.

Per esempio, se la tolleranza è `0.01`, il metodo della bisezione continuerà a dimezzare l'intervallo finché la differenza tra il limite superiore e quello inferiore è minore o uguale a `0.01`.

In questo laboratorio, implementerai una funzione che usa il metodo della bisezione per trovare la radice quadrata di un numero.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi definire una funzione chiamata `square_root_bisection` con tre parametri:
    - Il numero di cui vuoi trovare la radice quadrata.
    - La tolleranza, che è il margine di errore accettabile per il risultato. Devi impostare un valore di tolleranza predefinito.
    - Il numero massimo di iterazioni da eseguire. Devi impostare un numero predefinito di iterazioni.

1. La funzione `square_root_bisection` deve:
    - Sollevare un `ValueError` con il messaggio `Square root of negative number is not defined in real numbers` se il numero passato alla funzione è negativo.
    - Per i numeri `0` e `1`, stampare il messaggio: `The square root of [number] is [number]` e restituire il numero stesso come radice quadrata.
    - Per qualsiasi altro numero positivo, stampare la radice quadrata approssimata con il messaggio: `The square root of [square_target] is approximately [root]` e restituire il valore della radice calcolata.
    - Se nessun valore soddisfa la condizione di tolleranza, stampare un messaggio di fallimento: `Failed to converge within [maximum] iterations` e restituire `None`.

**Nota**: Non puoi importare alcun modulo per questo laboratorio.

# --hints--

Non devi importare alcun modulo.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

Devi avere una funzione chiamata `square_root_bisection`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

La tua funzione `square_root_bisection` deve avere tre parametri.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

Devi impostare un valore predefinito per la tolleranza e per il numero massimo di iterazioni.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

La tua funzione `square_root_bisection` deve sollevare un `ValueError` con il messaggio `Square root of negative number is not defined in real numbers` quando il numero passato alla funzione è negativo.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` deve restituire `0`.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` deve stampare `The square root of 0 is 0`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` deve restituire un numero tra `0.03162267660168379` e `0.031622876601683794`.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` deve stampare `The square root of 0.001 is approximately X`, dove `X` è un numero tra `0.03162267660168379` e `0.031622876601683794`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` deve restituire un numero tra `0.4999999` e `0.5000001`.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` deve stampare `The square root of 0.25 is approximately X`, dove `X` è un numero tra `0.4999999` e `0.5000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` deve restituire `1`.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` deve stampare `The square root of 1 is 1`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` deve restituire un numero tra `8.999` e `9.001`.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` deve stampare `The square root of 81 is approximately X`, dove `X` è un numero tra `8.999` e `9.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` deve restituire un numero tra `14.999` e `15.001`.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` deve stampare `The square root of 225 is approximately X`, dove `X` è un numero tra `14.999` e `15.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` deve restituire un numero tra `14.99999` e `15.00001`.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` deve stampare `The square root of 225 is approximately X`, dove `X` è un numero tra `14.99999` e `15.00001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` deve restituire un numero tra `14.9999999` e `15.0000001`.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` deve stampare `The square root of 225 is approximately X`, dove `X` è un numero tra `14.9999999` e `15.0000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` deve restituire `None`.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` deve stampare `Failed to converge within 10 iterations`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
