---
id: 68b94f23f7b3ee14078c89c5
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви використаєте елемент `iframe` для відображення відео. Для вас підготовлено базовий шаблонний код HTML.

Почніть зі створення елемента `h1` з текстом `iframe Video Display`.

# --hints--

У вас має бути елемент `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Ваш елемент `h1` має містити текст `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
