---
id: 68ba9c4f1688914d72876096
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Один із атрибутів — це `allow`. Це як список дозволів, який повідомляє браузеру, якими функціями `iframe` дозволено користуватися.

Ось елемент `iframe` з атрибутом `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Додайте атрибут `allow` зі значеннями `accelerometer`, `autoplay` і `clipboard-write`.

`accelerometer` дозволяє `iframe` використовувати датчики руху, щоб визначати такі речі, як нахил і обертання пристрою. `autoplay` дозволяє відео починати відтворення автоматично, а `clipboard-write` дозволяє iframe записувати дані у буфер обміну користувача.

# --hints--

Ваш елемент `iframe` має мати атрибут `allow` зі значенням `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

Атрибут `allow` вашого елемента `iframe` має містити `autoplay` як одне зі значень.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

Атрибут `allow` вашого елемента `iframe` має містити `clipboard-write` як одне зі значень.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
