---
id: 68ba9c4f1688914d72876095
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Елемент `iframe` також має атрибут `src` зі значенням, яке вказує URL або шлях до ресурсу для відображення.

Додайте атрибут `src` зі значенням `https://www.youtube.com/embed/I0_951_MPE0` до вашого елемента `iframe`.

На цьому етапі ви повинні бачити відео на сторінці, але є ще кілька атрибутів, які потрібно додати.

# --hints--

Ваш елемент `iframe` має мати атрибут `src`, встановлений у `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
