---
id: 68ba9c4f1688914d72876094
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

У першому уроці про елемент `iframe` ви дізналися, що це замінений елемент, як і `img`. Це означає, що він також може мати проперті `width` і `height`, які визначають його висоту та ширину.

Задайте вашому елементу `iframe` проперті `width` зі значенням `560` і проперті `height` зі значенням `315`.

# --hints--

Ваш елемент `iframe` має мати атрибут `width`, встановлений у `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Ваш елемент `iframe` має мати атрибут `height`, встановлений у `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
