---
id: 68ba9c4f1688914d72876097
title: Крок 6
challengeType: 0
dashedName: step-6
---

# --description--

Додайте `encrypted-media`, `gyroscope` і `web-share` до існуючих значень у атрибуті `allow`.

Ці три значення дозволять використовувати розширення для шифрування медіаконтенту, захищати відео, надавати доступ до датчиків руху та орієнтації пристрою, а також дозволять ділитися вмістом iframe через вбудовані діалоги спільного доступу пристрою.

# --hints--

Атрибут `allow` вашого елемента `iframe` повинен містити `encrypted-media` як одне зі своїх значень.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

Атрибут `allow` вашого елемента `iframe` повинен містити `gyroscope` як одне зі своїх значень.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

Атрибут `allow` вашого елемента `iframe` повинен містити `web-share` як одне зі своїх значень.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
