---
id: 612e6afc009b450a437940a1
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Empieza con la estructura básica de HTML. Agrega una declaracion `DOCTYPE` y los elementos `html`, `head`, `body` y `title`.

Establezca el idioma de esta página a inglés. Establece el `title` en `Piano`.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Tu elemento `html` debe tener una etiqueta de apertura con el atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Debes tener solo una etiqueta de apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Debes tener una etiqueta de cierre `head` tag.

```js
assert(code.match(/<\/head\s*>/i));
```

Debes tener una etiqueta de apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Debes tener una etiqueta de cierre `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Las etiquetas `head` y `body` deben ser hermanas.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

La etiqueta `head` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

La etiqueta `body` debe estar dentro de la etiqueta `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu proyecto debe tener el título `Piano`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'piano')
```

Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Piano');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
