---
id: 6140d10d50636e14695013b2
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

創建一個選擇器以標的你的第二個 `.line` 元素。將 `transform` 屬性設為 `rotate(60deg)`。

請記住，`transform` 屬性允許你操作元素的形狀。在這種情況下，使用 `rotate(60deg)` 值會使元素繞其 `transform-origin` 點順時針旋轉 60 度。

# --hints--

你應該創建一個 `.line:nth-of-type(2)` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

你的 `.line:nth-of-type(2)` 選擇器應該有一個設定為 `rotate(60deg)` 的 `transform` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform === 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
