---
id: 6140c7e645d8e905819f1dd4
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

從標準樣板開始。為你的文件添加 `DOCTYPE` 宣告、語言設定為英文的 `html` 元素、`head` 和 `body` 元素。

為你的 `meta` 元素添加正確的 `charset`，為你的 `title` 元素添加內容，並為 `./styles.css` 檔案添加一個 `link` 元素。

將 `title` 設定為 `Ferris Wheel`。

# --hints--

你的程式碼應該包含 `DOCTYPE` 宣告式。

```js
assert.match(code, /<!DOCTYPE/gi);
```

你應該在 `DOCTYPE` 宣告後面加上一個空格。

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

你應該定義文件型別為 `html`。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

你的 `html` 元素應該有一個開啟標籤，並且帶有 `lang` 屬性且值為 `en`。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你的 `html` 元素應該有一個結束標籤。

```js
assert.match(code, /<\/html\s*>/);
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

你應該有一個開啟的 `head` 頁籤。

```js
assert.match(code, /<head\s*>/i);
```

你應該有一個結束的 `head` 頁籤。

```js
assert.match(code, /<\/head\s*>/i);
```

你應該有一個開啟的 `body` 頁籤。

```js
assert.match(code, /<body\s*>/i);
```

你應該有一個結束的 `body` 頁籤。

```js
assert.match(code, /<\/body\s*>/i);
```

你的 `body` 元素應該是 `head` 元素之後的第一個元素。

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

你的 `head` 和 `body` 元素應該是同級元素。

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

你的 `head` 元素應該是你 `html` 元素內的第一個元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

你的 `head` 元素應該在 `html` 元素內。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

你的 `body` 元素應該位於 `html` 元素內。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

你的程式碼應該有一個 `meta` 元素。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

你的 `meta` 元素應該有一個值為 `UTF-8` 的 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

你的 `meta` 元素應該在 `head` 元素內。

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的專案應該有一個標題為 `Ferris Wheel`。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

你的程式碼應該有一個 `link` 元素。

```js
assert.match(code, /<link/);
```

你應該有一個 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
