---
id: 6140cd32d018ed0f600eefce
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

將 `.wheel` 選擇器的 `position` 屬性設為 `absolute`。將 `height` 和 `width` 都設為 `55vw`。

# --hints--

你的 `.wheel` 選擇器應該有一個設定為 `absolute` 的 `position` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.position === 'absolute');
```

你的 `.wheel` 選擇器應該設定 `height` 屬性為 `55vw`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.height === '55vw');
```

你的 `.wheel` 選擇器應該設定 `width` 屬性為 `55vw`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.width === '55vw');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
}
--fcc-editable-region--
```
