---
id: 617ace7d831f9c16a569b38a
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

為你的 `.wheel` 選擇器添加 `max-height` 和 `max-width` 屬性，兩者皆設定為 `500px`。

# --hints--

你的 `.wheel` 選擇器應該設定 `max-height` 屬性為 `500px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxHeight === '500px');
```

你的 `.wheel` 選擇器應該設定 `max-width` 屬性為 `500px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.maxWidth === '500px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
}
--fcc-editable-region--
```
