---
id: 616971b3cd990186b66c99a1
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Da Ihre externe CSS-Datei jetzt eingerichtet ist, können Sie mit dem Gestalten der Seite beginnen.

Zur Erinnerung: So sprechen Sie ein Absatz-Element an und richten es rechtsbündig aus:

```css
p {
  text-align: right;
}
```

Erstellen Sie eine neue CSS-Regel, die das `h1`-Element anspricht, und setzen Sie dessen `text-align`-Eigenschaft auf `center`.

# --hints--

Sie sollten einen `h1`-Selektor verwenden.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('h1'));
```

Ihre `h1`-CSS-Regel sollte eine `text-align`-Eigenschaft mit dem Wert `center` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
