---
id: 616968c2c94c8071b349c146
title: Step 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Verschachteln Sie innerhalb des `body` ein `h1`-Element mit dem Text `CSS Color Markers`.

# --hints--

Ihr Code sollte ein `h1`-Element enthalten.

```js
const title = document.querySelector('h1');
assert.exists(title);
```

Sie sollten ein öffnendes `h1`-Tag haben.

```js
assert.match(code, /<h1\s*>/i);
```

Sie sollten ein schließendes `h1`-Tag haben.

```js
assert.match(code, /<\/h1\s*>/i);
```

Ihr `h1`-Element sollte sich innerhalb des `body`-Elements befinden.

```js
assert.strictEqual(document.querySelector('h1')?.parentElement?.localName, 'body');
```

Ihr `h1`-Element sollte den Text `CSS Color Markers` enthalten.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim()?.toLowerCase(), 'css color markers')
```

Denken Sie daran, dass Groß- und Kleinschreibung sowie Rechtschreibung für den `h1`-Text wichtig sind.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent?.trim(), 'CSS Color Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
