---
id: 61696ef7ac756c829f9e4048
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

In diesem Projekt arbeiten Sie mit einer externen CSS-Datei, um die Seite zu gestalten. Eine `styles.css` Datei wurde bereits für Sie erstellt. Bevor Sie sie verwenden können, müssen Sie sie jedoch mit der Seite verknüpfen.

Verschachteln Sie ein `link`-Element innerhalb des `head`-Elements. Geben Sie ihm ein `rel`-Attribut mit dem Wert `"stylesheet"` und ein `href`-Attribut mit dem Wert `"styles.css"`.

# --hints--

Sie sollten ein `link`-Element haben.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Ihr `link`-Element sollte ein Void-Element sein, es darf kein schließendes Tag `</link>` haben.

```js
assert.notMatch(code, /<\/link>/i);
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `"stylesheet"` haben.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel")?.toLowerCase();
assert.strictEqual(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `"styles.css"` haben.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

Ihr `link`-Element sollte sich innerhalb des `head`-Elements befinden.

```js
const linkElementRegex = /<head\s*>(?:.|\r|\n)*?<link(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, linkElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css

```
