---
id: 616d4a84b756d9c4b8255093
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

Es ist Zeit, dem Marker etwas Farbe zu geben. Denken Sie daran, dass eine Möglichkeit, einem Element Farbe hinzuzufügen, die Verwendung eines <dfn>Farb-Schlüsselworts</dfn> wie `black`, `cyan` oder `yellow` ist.

Zur Erinnerung: So sprechen Sie die Klasse `freecodecamp` an:

```css
.freecodecamp {
  
}
```

Erstellen Sie eine neue CSS-Regel, die die Klasse `marker` anspricht, und setzen Sie deren `background-color`-Eigenschaft auf `red`.

**Hinweis:** Nach dem Hinzufügen des CSS werden Sie keine Änderungen sehen.

# --hints--

Sie sollten einen Klassenselektor erstellen, um die Klasse `marker` anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.marker'));
```

Ihre `.marker`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `red` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
