---
id: 616d595270d902f0e7086e18
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

Obwohl Sie drei separate Marker-`div`-Elemente haben, sehen sie aus wie ein großes Rechteck. Sie sollten etwas Abstand zwischen ihnen hinzufügen, damit jedes Element leichter zu erkennen ist.

Wenn die Kurzschreibweise `margin` zwei Werte hat, setzt sie `margin-top` und `margin-bottom` auf den ersten Wert und `margin-left` und `margin-right` auf den zweiten Wert.

Setzen Sie in Ihrer `.marker`-CSS-Regel die `margin`-Eigenschaft auf `10px auto`.

# --hints--

Ihre `.marker`-CSS-Regel sollte eine `margin`-Eigenschaft haben, die auf `10px auto` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, '10px auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  margin: auto;
--fcc-editable-region--
}

```
