---
id: 616d525007b8c5d8b3308b1c
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Ihr Marker würde besser aussehen, wenn er auf der Seite zentriert wäre. Eine einfache Möglichkeit, dies zu erreichen, ist die Verwendung der `margin` <dfn>Kurzschreibweise-Eigenschaft</dfn>.

Sie können den Außenabstandsbereich von Elementen mit Eigenschaften wie `margin-top` und `margin-left` separat einstellen. Die `margin` Kurzschreibweise-Eigenschaft erleichtert es, mehrere Außenabstandsbereiche gleichzeitig zu setzen.

Um Ihren Marker auf der Seite zu zentrieren, setzen Sie seine `margin`-Eigenschaft auf `auto`. Dadurch werden `margin-top`, `margin-right`, `margin-bottom` und `margin-left` alle auf `auto` gesetzt.

# --hints--

Ihre `.marker`-CSS-Regel sollte eine `margin`-Eigenschaft mit dem Wert `auto` enthalten.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.marker')?.margin, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
--fcc-editable-region--
  
--fcc-editable-region--
}

```
