---
id: 616d55bd160a95e22453a081
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Jetzt, wo Sie einen Marker mit Farbe zentriert haben, ist es Zeit, die anderen Marker hinzuzufügen.

Fügen Sie im `.container`-Element zwei weitere `div`-Elemente hinzu und geben Sie ihnen jeweils eine Klasse von `marker`.

# --hints--

Ihr erstes neues `div`-Element sollte ein öffnendes Tag haben.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][2]);
```

Ihr erstes neues `div`-Element sollte ein schließendes Tag haben.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][2]);
```

Ihr zweites neues `div`-Element sollte ein öffnendes Tag haben.

```js
assert.exists([...code.matchAll(/<div.*?>/gi)][3]);
```

Ihr zweites neues `div`-Element sollte ein schließendes Tag haben.

```js
assert.exists([...code.matchAll(/<\/div\s*>/gi)][3]);
```

Ihre neuen `div`-Elemente sollten sich innerhalb des `div` mit der Klasse `container` befinden.

```js
const containerChildren = [...document.querySelector('.container')?.children];
containerChildren.forEach(child => assert.strictEqual(child?.localName, 'div'));
assert.lengthOf(containerChildren, 3);
```

Ihre neuen `div`-Elemente sollten beide eine Klasse von `marker` haben.

```js
const containerChildren = [...document.querySelector('.container')?.children];
assert.isNotEmpty(containerChildren);
containerChildren.forEach(child => assert.isTrue(child?.classList?.contains('marker')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: auto;
}

```
