---
id: 6176519636a76b810ab1219a
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

Erstellen Sie eine CSS-Regel, die die Klasse `two` anspricht, und setzen Sie deren `background-color`-Eigenschaft auf `green`.

Erstellen Sie außerdem eine separate CSS-Regel, die die Klasse `three` anspricht, und setzen Sie deren `background-color` auf `blue`.

# --hints--

Sie sollten einen Klassenselektor verwenden, um die Klasse `two` anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.two'));
```

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `green` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Sie sollten einen Klassenselektor verwenden, um die Klasse `three` anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.three'));
```

Ihre `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `blue` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
