---
id: 61764c602bee6974e7790f35
title: Schritt 11
challengeType: 0
dashedName: step-11
---

# --description--

Um den Markierungen unterschiedliche Farben zu geben, müssen Sie jeder eine eindeutige Klasse hinzufügen. Sie können einem Element mehrere Klassen hinzufügen, indem Sie diese im `class`-Attribut auflisten und durch ein Leerzeichen trennen. Zum Beispiel fügt der folgende Code sowohl die Klassen `animal` als auch `dog` zu einem `div`-Element hinzu:

```html
<div class="animal dog">
```

Wenn Sie einem HTML-Element mehrere Klassen hinzufügen, können die Styles der zuerst aufgelisteten Klassen im Stylesheet durch später folgende Klassen überschrieben werden.

Beginnen Sie damit, die Klasse `one` zum ersten Marker-`div`-Element hinzuzufügen.


# --hints--

Sie sollten die Klasse `one` zum ersten Marker-`div`-Element hinzufügen.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

Ihr erster Marker-`div` sollte die Klassen `marker` und `one` haben.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert.isTrue(containerFirstChild?.classList?.contains('marker'));
assert.isTrue(containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
--fcc-editable-region--
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
