---
id: 617b8e0d93a8d10d9a90e720
title: Step 18
challengeType: 0
dashedName: step-18
---

# --description--

Beachten Sie, dass der `background-color` für Ihren Marker immer noch rot ist. Das liegt daran, dass Sie den roten Wert der `rgb`-Funktion auf das Maximum von `255` gesetzt haben, also 100 % Rot, und sowohl den grünen als auch den blauen Wert auf `0` gesetzt haben.

Verwenden Sie nun die `rgb`-Funktion, um die anderen Farben einzustellen.

Verwenden Sie in der `.two`-CSS-Regel die `rgb`-Funktion, um die `background-color` auf den Maximalwert für Grün zu setzen, und `0` für die anderen Werte. Und in der `.three`-CSS-Regel verwenden Sie die `rgb`-Funktion, um die `background-color` auf den Maximalwert für Blau zu setzen, und `0` für die anderen Werte.

# --hints--

Ihre `.two`-CSS-Regel sollte das Farb-Schlüsselwort `green` nicht verwenden, um die `background-color`-Eigenschaft zu setzen.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'green');
```

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `rgb(0, 255, 0)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

Ihre `.three`-CSS-Regel sollte das Farb-Schlüsselwort `blue` nicht verwenden, um die `background-color`-Eigenschaft zu setzen.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'blue');
```

Ihre `.three`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `rgb(0, 0, 255)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

--fcc-editable-region--
.two {
  background-color: green;
}

.three {
  background-color: blue;
}
--fcc-editable-region--

```
