---
id: 617b65579ce424bf5f02ca73
title: Schritt 16
challengeType: 0
dashedName: step-16
---

# --description--

Wie Sie in den Lektionen gelernt haben, gibt es zwei Hauptfarbmodelle: das <dfn>additive</dfn> RGB-Modell (Rot, Grün, Blau), das in elektronischen Geräten verwendet wird, und das <dfn>subtraktive</dfn> CMYK-Modell (Cyan, Magenta, Gelb, Schwarz), das im Druck eingesetzt wird.

In diesem Projekt arbeiten Sie mit dem RGB-Modell. Das bedeutet, Farben beginnen als Schwarz und ändern sich, wenn unterschiedliche Mengen an Rot, Grün und Blau hinzugefügt werden. Eine einfache Möglichkeit, dies zu sehen, ist die CSS-Funktion `rgb`.

Erstellen Sie eine neue CSS-Regel, die die Klasse `container` anspricht, und setzen Sie deren `background-color` mit `rgb(0, 0, 0)` auf Schwarz.

# --hints--

Sie sollten einen Klassenselektor verwenden, um die Klasse `container` anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.container'));
```

Ihre `.container`-CSS-Regel sollte eine `background-color`-Eigenschaft enthalten, die auf `rgb(0, 0, 0)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.container')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
