---
id: 617b92b9de349513466f6156
title: Step 19
challengeType: 0
dashedName: step-19
---

# --description--

Obwohl die roten und blauen Marker gleich aussehen, ist der grüne viel heller als zuvor. Das liegt daran, dass das `green` Farb-Schlüsselwort tatsächlich ein dunklerer Farbton ist und ungefähr zwischen Schwarz und dem Maximalwert für Grün liegt.

Setzen Sie in der `.two`-CSS-Regel den Grünwert in der `rgb`-Funktion auf `127`, um die Intensität zu verringern.

# --hints--

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(0, 127, 0)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 127, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
