---
id: 617b97abd9f3f61d1590b815
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

Im additiven RGB-Farbmodell sind <dfn>Primärfarben</dfn> Farben, die, wenn sie kombiniert werden, reines Weiß erzeugen. Damit dies geschieht, muss jede Farbe ihre höchste Intensität haben.

Bevor Sie Farben kombinieren, setzen Sie Ihren grünen Marker wieder auf reines Grün zurück. Für die `rgb`-Funktion in der `.two`-CSS-Regel setzen Sie Grün wieder auf den Maximalwert von `255`.

# --hints--

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `rgb(0, 255, 0)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 0, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 127, 0);
--fcc-editable-region--
}

.three {
  background-color: rgb(0, 0, 255);
}

```
