---
id: 617bc3386dc7d07d6469bf20
title: Schritt 27
challengeType: 0
dashedName: step-27
---

# --description--

Beachten Sie, dass Sie zur Erstellung von Orange die Intensität von Rot erhöht und die Intensität von Grün in den `rgb`-Werten verringert haben. Das liegt daran, dass Orange durch das Mischen von Rot mit Gelb entsteht.

Um die tertiäre Farbe Frühlingsgrün zu erzeugen, kombinieren Sie Cyan mit Grün. Aktualisieren Sie die `rgb`-Funktion in der `.two`-CSS-Regel, indem Sie Grün auf den Maximalwert setzen und Blau auf `127`.

# --hints--

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(0, 255, 127)` enthalten.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 127)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: rgb(255, 127, 0);
}

.two {
--fcc-editable-region--
  background-color: rgb(0, 255, 255);
--fcc-editable-region--
}

.three {
  background-color: rgb(255, 0, 255);
}

```
