---
id: 617b8b38f32bf2080a140675
title: Schritt 17
challengeType: 0
dashedName: step-17
---

# --description--

Wie Sie in vorherigen Lektionen gelernt haben, ist eine Funktion ein Codeabschnitt, der eine Eingabe annehmen und eine bestimmte Aktion ausführen kann. Die CSS-`rgb`-Funktion akzeptiert Werte oder <dfn>Argumente</dfn> für Rot, Grün und Blau und erzeugt eine Farbe:

```css
rgb(red, green, blue);
```

Jeder Rot-, Grün- und Blauwert ist eine Zahl von `0` bis `255`. `0` bedeutet, dass 0 % dieser Farbe vorhanden sind, und ergibt Schwarz. `255` bedeutet, dass 100 % dieser Farbe vorhanden sind.

Ersetzen Sie in der `.one`-CSS-Regel das Farb-Schlüsselwort `red` durch die `rgb`-Funktion. Für die `rgb`-Funktion setzen Sie den Wert für Rot auf `255`, den Wert für Grün auf `0` und den Wert für Blau auf `0`.

# --hints--

Ihre `.one`-CSS-Regel sollte das Farb-Schlüsselwort `red` nicht verwenden, um die Eigenschaft `background-color` zu setzen.

```js
assert.notStrictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

Ihre `.one`-CSS-Regel sollte eine Eigenschaft `background-color` mit dem Wert `rgb(255, 0, 0)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
--fcc-editable-region--
  background-color: red;
--fcc-editable-region--
}

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
