---
id: 618a0927005553b74bfa05ae
title: Step 31
challengeType: 0
dashedName: step-31
---

# --description--

Ein Farbrad ist ein Kreis, bei dem ähnliche Farben oder <dfn>Farbtöne</dfn> nahe beieinanderliegen und unterschiedliche weiter auseinander. Zum Beispiel liegt reines Rot zwischen den Farbtönen Rosa und Orange.

Zwei Farben, die sich auf dem Farbrad gegenüberliegen, nennt man <dfn>komplementäre Farben</dfn>. Wenn zwei komplementäre Farben kombiniert werden, erzeugen sie Grau. Werden sie jedoch nebeneinander platziert, erzeugen diese Farben einen starken visuellen Kontrast und wirken heller.

Setzen Sie in der `rgb`-Funktion für die `.one`-CSS-Regel den Rotwert auf das Maximum von `255`, um reines Rot zu erzeugen. Setzen Sie in der `rgb`-Funktion für die `.two`-CSS-Regel die Werte für Grün und Blau auf das Maximum von `255`, um Cyan zu erzeugen.

# --hints--

Ihre `.one`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(255, 0, 0)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'rgb(255, 0, 0)');
```

Ihre `.two`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(0, 255, 255)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor, 'rgb(0, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}
--fcc-editable-region--

.three {
  background-color: rgb(0, 0, 0);
}

```
