---
id: 619b72a0db211f1c29afb906
title: Step 37
challengeType: 0
dashedName: step-37
---

# --description--

Aktualisieren Sie die `.one`-CSS-Regel, um die neue `red`-Klasse anzusprechen.

# --hints--

Ihr Code sollte keine `.one`-Klassenselektor mehr enthalten.

```js
assert.notExists(new __helpers.CSSHelp(document).getStyle('.one'));
```

Sie sollten einen Klassenselektor verwenden, um die Klasse `red` anzusprechen.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.red'));
```

Ihre `.red`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `rgb(0, 0, 0)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.red')?.backgroundColor, 'rgb(0, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}
```
