---
id: 619b7fd56aa2253778bcf5f7
title: Step 43
challengeType: 0
dashedName: step-43
---

# --description--

Das <dfn>HSL</dfn>-Farbmodell, also Farbton, Sättigung und Helligkeit, ist eine weitere Möglichkeit, Farben darzustellen.

Die CSS-`hsl`-Funktion akzeptiert 3 Werte: eine Zahl von 0 bis 360 für den Farbton, einen Prozentsatz von 0 bis 100 für die Sättigung und einen Prozentsatz von 0 bis 100 für die Helligkeit.

Wenn Sie sich ein Farbrad vorstellen, liegt der Farbton Rot bei 0 Grad, Grün bei 120 Grad und Blau bei 240 Grad.

Die Sättigung ist die Intensität einer Farbe von 0 %, also Grau, bis 100 % für reine Farbe. Sie müssen das Prozentzeichen `%` bei den Werten für Sättigung und Helligkeit hinzufügen.

Die Helligkeit gibt an, wie hell eine Farbe erscheint, von 0 %, also komplett Schwarz, bis 100 %, komplett Weiß, wobei 50 % neutral ist.

Verwenden Sie in der `.blue`-CSS-Regel die `hsl`-Funktion, um die `background-color`-Eigenschaft auf reines Blau zu ändern. Setzen Sie den Farbton auf `240`, die Sättigung auf `100%` und die Helligkeit auf `50%`.

# --hints--

Ihre `.blue`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `hsl(240, 100%, 50%)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.blue')?.backgroundColor, 'rgb(0, 0, 255)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
  background-color: #007F00;
}

.blue {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

```
