---
id: 619b761916dac02643940022
title: Step 41
challengeType: 0
dashedName: step-41
---

# --description--

Eine sehr gebräuchliche Methode, um mit CSS Farbe auf ein Element anzuwenden, ist die Verwendung von <dfn>hexadezimalen</dfn> oder Hex-Werten. Obwohl Hex-Werte kompliziert klingen, sind sie eigentlich nur eine andere Form von RGB-Werten.

Hex-Farbwerte beginnen mit einem `#`-Zeichen und bestehen aus sechs Zeichen von 0-9 und A-F. Das erste Zeichenpaar steht für Rot, das zweite für Grün und das dritte für Blau. Zum Beispiel `#4B5320`.

Im `.green`-Klassen-Selektor setzen Sie die `background-color`-Eigenschaft auf einen Hex-Farbcode mit den Werten `00` für Rot, `FF` für Grün und `00` für Blau.

# --hints--

Ihre `.green`-CSS-Regel sollte eine `background-color`-Eigenschaft mit dem Wert `#00FF00` enthalten.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.green')?.backgroundColor, 'rgb(0, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.green {
--fcc-editable-region--
  background-color: rgb(0, 0, 0);
--fcc-editable-region--
}

.blue {
  background-color: rgb(0, 0, 0);
}

```
