---
id: 61a493ead935c148d2b76312
title: Schritt 45
challengeType: 0
dashedName: step-45
---

# --description--

Die `linear-gradient`-Funktion ist sehr flexibel – hier ist die grundlegende Syntax, die Sie in diesem Tutorial überprüfen werden:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` ist die Richtung der Linie, die für die Transition verwendet wird. `color1` und `color2` sind Farb-Argumente, also die Farben, die in der Transition selbst verwendet werden. Diese können jede Art von Farbe sein, einschließlich Farb-Schlüsselwörter, Hex, `rgb` oder `hsl`.

Jetzt wenden Sie einen Farbverlauf von Rot zu Grün entlang einer 90-Grad-Linie auf den ersten Marker an.

Setzen Sie zuerst in der `.red`-CSS-Regel die `background`-Eigenschaft auf `linear-gradient()` und übergeben Sie ihr den Wert `90deg` als `gradientDirection`.

# --hints--

Ihre `.red`-CSS-Regel sollte eine `background`-Eigenschaft mit dem Wert `linear-gradient(90deg)` enthalten.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
