---
id: 61a49d15bdbb5e57cc6fd280
title: Schritt 49
challengeType: 0
dashedName: step-49
---

# --description--

Farbstopps ermöglichen es Ihnen, genau festzulegen, wo Farben entlang der Verlaufslinie platziert werden. Sie werden in der `linear-gradient`-Funktion verwendet, um anzugeben, wo eine Farbe endet und der Übergang zur nächsten Farbe beginnt. Zur Definition der Position des Farbstopps kann entweder ein Prozentwert oder ein Längenwert verwendet werden.

Zum Beispiel findet bei diesem Rot-Schwarz-Verlauf der Übergang von Rot zu Schwarz am 90%-Punkt entlang der Verlaufslinie statt, sodass Rot den größten Teil des verfügbaren Bereichs einnimmt:

```css
linear-gradient(90deg, red 90%, black);
```

Fügen Sie in der `linear-gradient`-Funktion nach dem ersten roten Farbargument einen `75%`-Farbstopp hinzu. Fügen Sie keine Farbstopps zu den anderen Farbargumenten hinzu.

# --hints--

Ihre `.red`-CSS-Regel sollte eine `background`-Eigenschaft mit dem Wert `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))` haben.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
