---
id: 61a489b8579e87364b2d2cdb
title: Schritt 44
challengeType: 0
dashedName: step-44
---

# --description--

Sie haben bereits einige Möglichkeiten geübt, um einfarbige Farben in CSS festzulegen, aber Sie können auch einen Farbverlauf, also ein <dfn>Gradient</dfn>, auf ein Element anwenden.

Ein Gradient ist eine Übergangsfärbung, bei der eine Farbe in eine andere übergeht. Die CSS-`linear-gradient`-Funktion ermöglicht es Ihnen, die Richtung des Übergangs entlang einer Linie zu steuern und festzulegen, welche Farben verwendet werden.

Wichtig ist, dass die `linear-gradient`-Funktion tatsächlich ein `image`-Element erzeugt und normalerweise mit der `background`-Eigenschaft kombiniert wird, die als Wert ein Bild akzeptieren kann.

Ändern Sie in der `.red`-CSS-Regel die `background-color`-Eigenschaft in `background`.

# --hints--

Ihre `.red`-CSS-Regel sollte eine `background`-Eigenschaft mit dem Wert `rgb(255, 0, 0)` enthalten.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:rgb\(255,0,0\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
--fcc-editable-region--
  background-color: rgb(255, 0, 0);
--fcc-editable-region--
}

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
