---
id: 61b093329e7fc020b43b1bb5
title: Schritt 70
challengeType: 0
dashedName: step-70
---

# --description--

Sie haben die `rgb`-Funktion verwendet, um Farben festzulegen. Um einem `rgb`-Farbwert einen Alphakanal hinzuzufügen, verwenden Sie stattdessen die `rgba`-Funktion.

Die `rgba`-Funktion funktioniert genauso wie die `rgb`-Funktion, nimmt aber eine weitere Zahl von `0` bis `1.0` für den Alphakanal entgegen:

```css
rgba(redValue, greenValue, blueValue, alphaValue);
```

Sie können auch bei `hsl`- und `hex`-Farben einen Alphakanal verwenden. Sie werden bald sehen, wie das geht.

In der `.sleeve`-Regel verwenden Sie die `rgba`-Funktion, um die `background-color`-Eigenschaft auf reines Weiß mit 50 % Deckkraft zu setzen.

# --hints--

Ihre `.sleeve`-CSS-Regel sollte eine `background-color`-Eigenschaft haben, die auf `rgba(255, 255, 255, 0.5)` gesetzt ist.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.sleeve')?.backgroundColor, 'rgba(255, 255, 255, 0.5)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.sleeve {
  width: 110px;
  height: 25px;
--fcc-editable-region--
  background-color: white;
--fcc-editable-region--
}

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
