---
id: 6690e10ebe2181212abc9652
title: Progetta un biglietto da visita
challengeType: 25
dashedName: design-a-business-card
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un `div` con un attributo `class` con valore `business-card` che conterrà tutti gli altri elementi.  
1. All'interno dell'elemento `.business-card` ci dovrebbe essere un elemento `img` con un attributo `class` con valore `profile-image`. Puoi impostare la sorgente dell'immagine su `https://cdn.freecodecamp.org/curriculum/labs/flower.jpg` se vuoi. Ci dovrebbe essere un `alt` con una descrizione significativa.  
1. All'interno dell'elemento `.business-card`, dopo l'elemento `img`, dovresti avere tre elementi `p` con un attributo `class` che ha, rispettivamente, valore `full-name`, `designation` e `company`.  
1. Il primo elemento `p` dovrebbe contenere il tuo nome.  
1. Il secondo elemento `p` dovrebbe contenere la tua qualifica.  
1. Il terzo elemento `p` dovrebbe contenere il nome della tua azienda.  
1. Ci dovrebbe essere un elemento `hr` sotto il terzo elemento `p`.  
1. Dopo l'elemento `hr` ci dovrebbero essere due elementi `p`. Nel primo elemento `p` dovresti avere un indirizzo email come testo. Nel secondo elemento `p` il testo dovrebbe essere un numero di telefono.  
1. Dopo i due elementi `p` ci dovrebbe essere un elemento `a` con classe `portfolio-link`, testo `Portfolio` e dovrebbe collegarsi a un URL valido.  
1. Ci dovrebbe essere un elemento `hr` dopo l'elemento `a` contenente il collegamento al portfolio.  
1. Dovresti avere un altro elemento `div` con un attributo `class` con valore `social-media`. All'interno di questo elemento ci dovrebbe essere un elemento `h2` con il testo `Connect with me`.  
1. All'interno degli elementi `.social-media` ci dovrebbero essere tre elementi `a`. Gli elementi `a` dovrebbero avere rispettivamente i testi `Twitter`, `LinkedIn` e `GitHub` con collegamenti a siti web validi.  
1. Dovresti collegare il foglio di stile `styles.css` usando il tag `link` nella sezione `head` del file HTML.  
1. Lo sfondo della pagina dovrebbe essere `rosybrown`. Il carattere tipografico generale dovrebbe essere `Arial`, con fallback `sans-serif`.  
1. Il selettore `.business-card` dovrebbe avere proprietà per impostare l'elemento con larghezza `300px` e un colore di sfondo a tua scelta. Dovresti anche impostare un margine `padding` di `20px` tutt'intorno e un margine superiore `margin` di `100px`. Il testo dovrebbe essere allineato al centro e la dimensione del carattere dovrebbe essere `16px`. I margini sinistro e destro dovrebbero essere impostati a `auto`.  
1. Il selettore `.profile-image` dovrebbe avere una proprietà `max-width` con valore `100%`.  
1. Tutti gli elementi `p` dovrebbero avere un margine superiore e inferiore `margin` di `5px`.  
1. Tutti i collegamenti nella pagina non dovrebbero essere sottolineati.

# --hints--

Dovresti avere un elemento principale `div` con un attributo `class` con valore `business-card`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'business-card');
```

All'interno dell'elemento `.business-card` ci dovrebbe essere un elemento `img` con un attributo `class` con valore `profile-image`.

```js
assert.exists(document.querySelector('img.profile-image'));
```

La sorgente dell'immagine dovrebbe essere impostata su un'immagine valida.

```js
assert.isAbove(document.querySelector('.profile-image')?.getAttribute('src')?.length, 0);
```

L'attributo `alt` dell'immagine dovrebbe essere impostato su una descrizione significativa.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

All'interno dell'elemento `.business-card`, dopo l'elemento `img`, ci dovrebbe essere un elemento `p` con un attributo `class` con valore `full-name`.

```js
assert.exists(document.querySelector('div.business-card p.full-name'));
```

Dopo l'elemento `.full-name` ci dovrebbe essere un elemento `p` con un attributo `class` con valore `designation`.

```js
assert.exists(document.querySelector('div.business-card p.designation'));
```

Dopo l'elemento `.designation` ci dovrebbe essere un elemento `p` con un attributo `class` con valore `company`.

```js
assert.exists(document.querySelector('div.business-card p.company'));
```

Il primo elemento `p`, quello con classe `full-name`, dovrebbe contenere il tuo nome.

```js
assert.isNotEmpty(document.querySelector(".full-name")?.innerText.trim());
```

Il secondo elemento `p`, quello con un attributo `class` di valore `designation`, dovrebbe contenere la tua qualifica.

```js
assert.isNotEmpty(document.querySelector(".designation")?.innerText.trim());
```

Il terzo elemento `p`, quello con l'attributo `class` con valore `company`, dovrebbe contenere il nome della tua azienda.

```js
assert.isNotEmpty(document.querySelector(".company")?.innerText.trim());
```

Dopo il terzo elemento `p`, quello con l'attributo `class` con valore `company`, ci dovrebbe essere un elemento `hr`.

```js
assert.exists(document.querySelector('p.company + hr'));
```

Dopo il primo elemento `hr` ci dovrebbe essere un elemento `p` con un indirizzo email come testo.

```js
assert.isNotEmpty(document.querySelectorAll('.business-card hr + p')[0].innerText.trim());
```

Dopo l'elemento `p` con l'email, ci dovrebbe essere un altro elemento `p` con un numero di telefono come testo.

```js
assert.match(document.querySelectorAll(".business-card p")[4].textContent,/.+/);
```

Dopo l'elemento `p` con il numero di telefono, ci dovrebbe essere un elemento `a` con classe `portfolio-link`.

```js
assert.exists(document.querySelector('a.portfolio-link'));
```

L'elemento `.portfolio-link` dovrebbe avere il testo `Portfolio`.

```js
assert.equal(document.querySelector('a.portfolio-link').innerText.trim(), 'Portfolio');
```

Il tuo elemento `.portfolio-link` dovrebbe avere un `href` valido.

```js
assert.isAbove(document.querySelector('.portfolio-link').getAttribute('href').length, 0);
```

Dopo l'elemento `a` contenente il portfolio, ci dovrebbe essere un altro elemento `hr` come divisore.

```js
assert.exists(document.querySelector('a + hr'));
```

Dopo il secondo elemento `hr` ci dovrebbe essere un elemento `div` con un attributo `class` con valore `social-media`.

```js
assert.exists(document.querySelector('div.business-card div.social-media'));
```

All'interno dell'elemento `.social-media` ci dovrebbe essere un elemento `h2` con il testo `Connect with me`.

```js
assert.equal(document.querySelector('.social-media h2')?.textContent.trim(), 'Connect with me');
```

All'interno dell'elemento `.social-media` ci dovrebbero essere tre elementi `a`.

```js
assert.lengthOf(document.querySelectorAll('.social-media a'), 3);
```

Il primo elemento `a` dovrebbe avere il testo `Twitter`.

```js
assert.equal(document.querySelectorAll(".social-media a")[0]?.textContent.trim(), 'Twitter');
```

L'attributo `href` del primo elemento `a` dovrebbe puntare a un collegamento valido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[0]?.getAttribute('href')?.length, 0);
```

Il secondo elemento `a` dovrebbe avere il testo `LinkedIn`.

```js

assert.equal(document.querySelectorAll(".social-media a")[1]?.textContent.trim(), 'LinkedIn');
```

L'attributo `href` del secondo elemento `a` dovrebbe puntare a un collegamento valido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[1]?.getAttribute('href')?.length, 0);
```

Il terzo elemento `a` dovrebbe avere il testo `GitHub`.

```js
assert.equal(document.querySelectorAll(".social-media a")[2]?.textContent.trim(), 'GitHub');
```

L'attributo `href` del terzo elemento `a` dovrebbe puntare a un collegamento valido.

```js
assert.isAbove(document.querySelectorAll('.social-media a')[2]?.getAttribute('href')?.length, 0);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');
```

Dovresti impostare il colore di sfondo della pagina nel selettore `body` su `rosybrown`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('background-color'), 'rosybrown');
```

Dovresti impostare il carattere tipografico della pagina su `Arial` nell'elemento `body` con fallback `sans-serif`.

```js
const el = new __helpers.CSSHelp(document).getStyle('body')?.getPropVal('font-family');
const condition1= el === 'Arial, sans-serif';
const condition2= el === '"Arial", sans-serif';
assert.isTrue(condition1 || condition2);
```

Il tuo selettore `.business-card` dovrebbe avere una proprietà `width` con valore `300px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('width'), '300px');
```

Il tuo selettore `.business-card` dovrebbe avere una proprietà `background-color` con un valore colore valido.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.business-card')?.backgroundColor);
```

Il tuo selettore `.business-card` dovrebbe avere una proprietà `padding` con valore `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('padding'), '20px');
```

Il tuo selettore `.business-card` dovrebbe avere un margine superiore di `100px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-top'), '100px');
```

Il testo all'interno dell'elemento `.business-card` dovrebbe essere allineato al centro.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('text-align'), 'center');
```

La dimensione del carattere del testo all'interno dell'elemento `.business-card` dovrebbe essere `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('font-size'), '16px');
```

I margini sinistro e destro dell'elemento `.business-card` dovrebbero essere impostati a `auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-left'), 'auto');
assert.equal(new __helpers.CSSHelp(document).getStyle('.business-card')?.getPropVal('margin-right'), 'auto');
```

Il tuo selettore `.profile-image` dovrebbe avere una proprietà `max-width` con valore `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.profile-image')?.getPropVal('max-width'), '100%');
```

I tuoi elementi `p` dovrebbero avere un margine superiore e inferiore `margin` di `5px`.

```js
const marginTop = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-top');
const marginBottom = new __helpers.CSSHelp(document).getStyle('p')?.getPropVal('margin-bottom');

assert.equal(marginTop, '5px');
assert.equal(marginBottom, '5px');
```

I collegamenti della pagina non dovrebbero avere sottolineatura.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('a')?.getPropVal('text-decoration'), 'none');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Business Card</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Digital Business Card</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="business-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/flower.jpg" alt="A vibrant flower in full bloom" class="profile-image">
        <p class="full-name">Your Name</p>
        <p class="designation">Software Developer</p>
        <p class="company" >@freeCodeCamp</p>
        <hr>
        <p>Email: your-email@example.com</p>
        <p>Phone: (123) 456-7890</p>
        <a class="portfolio-link" href="http://example.com">Portfolio</a>
        <hr>
        <div class="social-media">
            <h2>Connect with me</h2>
            <a href="http://twitter.com/">Twitter</a>
            <a href="http://linkedin.com/">LinkedIn</a>
            <a href="http://github.com/">GitHub</a>
        </div>
    </div>
</body>

</html>

```

```css
body {
    background-color: rosybrown;
    font-family: 'Arial', sans-serif;
}

.business-card {
    background-color: lavenderblush;
    border: 3px solid darkslategray;
    padding: 20px;
    text-align: center;
    font-size: 16px;
    width: 300px;
    margin-left: auto;
    margin-right: auto;
    margin-top: 100px;
}

.profile-image {
    max-width: 100%;
    width: 150px;
    height: 150px;
    margin-bottom: 10px;
    border: 2px solid lightgray;
}

.full-name {
    font-size: 24px;
    color: darkslategray;
}

.designation {
    font-size: 18px;
    color: gray;
}

.company {
    font-size: 16px;
    color: darkgray;
}

p {
    margin: 5px 0;
}

a {
    text-decoration: none;
}

.social-media {
    margin-top: 10px;
    font-size: 14px;
}

.social-media a {
    margin: 10px;
    color: maroon
}

hr {
    margin: 20px 0;
}
```
