---
id: 61537bb9b1a29430ac15ad38
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Within your `head` element, add a `title` element with the text set to `Photo Gallery`, and a `link` element to add your `styles.css` file to the page.

# --hints--

Dein `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

Dein Code sollte ein `title`-Element haben.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dein Projekt sollte einen Titel namens `Photo Gallery` tragen.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

Denke daran, dass Schreibweise und Rechtschreibung für den Titel wichtig sind.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
