---
id: 6998b99a8e195cd0c1f211df
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、JavaScriptオブジェクトを使ってシンプルな野生動物トラッカーを作成します。

こちらはオブジェクトの例です：

```js
const animal = {}; // empty object
```

これは空のオブジェクトを作成します。

次に、`tiger` という名前の変数を作成し、空のオブジェクトを代入してください。

# --hints--

`tiger` という名前の変数を作成してください。

```js
assert.exists(tiger);
```

`tiger` 変数には空のオブジェクトを代入する必要があります。

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
