---
id: 6999b0d79ea938edcb032237
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

Wildlife Trackerワークショップのこの最終ステップでは、ブラケット記法を使ってオブジェクトのプロパティにアクセスする関数を作成します。

ブラケット記法を使ってプロパティにアクセスする例を示します：

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

この例では、ブラケット記法を使うことで変数を使ってプロパティにアクセスできます。

では、`getProperty`という名前の関数を作成してください。

この関数は2つのパラメータ、`animal`と`propertyName`を受け取ります。

ブラケット記法を使ってプロパティの値を返してください。

関数を作成したら、`console.log`を使って`getProperty(tiger, "species")`と`getProperty(elephant, "age")`を呼び出し、コンソールに返された値を表示してください。

これでWildlife Trackerワークショップは完了です！

# --hints--

`getProperty`という名前の関数を作成してください。

```js
assert.isFunction(getProperty);
```

`getProperty`関数は2つのパラメータ、`animal`と`propertyName`を持つ必要があります。

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty`はブラケット記法を使ってプロパティにアクセスしてください。

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

`getProperty(tiger, "species")`をコンソールにログ出力してください。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

`getProperty(elephant, "age")`をコンソールにログ出力してください。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

`getProperty(tiger, "species")`を呼び出すと`"Tiger"`が返されるはずです。

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

`getProperty(elephant, "age")`を呼び出すと`12`が返されるはずです。

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty`は関数のパラメータを使い、どんなオブジェクトでも動作するようにしてください。

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
