---
id: 6999b07c34c2cde21ee6e1d8
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

このステップでは、オブジェクトに特定のプロパティがあるかどうかをチェックする関数を作成します。

以下は `hasOwnProperty` を使った例です。

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

この例では、`hasColor` 関数が `hasOwnProperty` を使ってオブジェクトに `color` プロパティが存在するかどうかを確認しています。

では、`hasHabitat` という名前の関数を作成してください。この関数は `animal` という名前のパラメータを1つ受け取ります。

`hasOwnProperty` メソッドを使って、`true` オブジェクトに `animal` というプロパティがあれば `"habitat"` を返し、なければ `false` を返すようにしてください。

関数を作成したら、`console.log` を使って `hasHabitat(tiger)` と `hasHabitat(elephant)` を呼び出し、両方の結果をコンソールで確認してください。

# --hints--

`hasHabitat` という名前の関数を作成してください。

```js
assert.isFunction(hasHabitat);
```

`hasHabitat` 関数は `animal` という単一のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` は `hasOwnProperty` メソッドを使って `habitat` プロパティをチェックするべきです。

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

`hasHabitat(tiger)` をコンソールにログ出力してください。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

`hasHabitat(elephant)` をコンソールにログ出力してください。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

`hasHabitat(tiger)` を呼び出すと、生息地を追加した後に `true` を返すはずです。

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

生息地が存在しない場合、`hasHabitat(elephant)` を呼び出すと `false` を返すはずです。

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` は関数のパラメータを使い、どんなオブジェクトでも動作するようにしてください。

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
