---
id: 69999ff7f7862fdfb885ea5d
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

`tiger` オブジェクトに `age` という新しいプロパティを追加して更新してください。

`age` プロパティを `5` の値に設定してください。

# --hints--

`tiger` には `age` プロパティが必要です。

```js
assert.property(tiger, "age");
```

`tiger.age` は `5` と等しくする必要があります。

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
