---
id: 6999b0198d18d402c36cc5fe
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

このステップでは、オブジェクトからプロパティを削除する関数を作成します。

以下は `delete` を使ってプロパティを削除する例です。

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

この例では、`canFly` オブジェクトから `bird` プロパティが削除されています。

では、`removeEndangeredStatus` という名前の関数を作成してください。この関数は `animal` という名前のパラメータを1つ受け取ります。

関数の中で、`isEndangered` キーワードを使って `animal` オブジェクトから `delete` プロパティを削除してください。更新された `animal` オブジェクトを返します。

関数を作成したら、`console.log` を使って `removeEndangeredStatus(tiger)` を呼び出し、更新されたオブジェクトをコンソールで確認してください。

# --hints--

`removeEndangeredStatus` という名前の関数を作成してください。

```js
assert.isFunction(removeEndangeredStatus);
```

`removeEndangeredStatus` 関数は `animal` という単一のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`removeEndangeredStatus` 関数は更新された `animal` オブジェクトを返すべきです。

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` は `delete` キーワードを使って `isEndangered` プロパティを削除するべきです。

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

`removeEndangeredStatus(tiger)` をコンソールにログ出力してください。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

`removeEndangeredStatus(tiger)` を呼び出すと `isEndangered` プロパティが削除されるはずです。

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` は関数のパラメータを使い、どんなオブジェクトでも扱えるようにしてください。

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
