---
id: 6999ac171928136168654774
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

このステップでは、オブジェクトの`species`プロパティにアクセスする関数を作成します。

オブジェクトからプロパティにアクセスする小さな例を示します：

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

この例では、関数がオブジェクトをパラメータとして受け取り、その`species`プロパティを返します。

次に、`getSpecies`という名前の関数を作成してください。

関数は`animal`という名前のパラメータを1つ受け取る必要があります。

ドット記法を使って`species`の`animal`プロパティを返してください。

関数を作成したら、`console.log`を使って`getSpecies(tiger)`を呼び出し、コンソールで結果を確認してください。

# --hints--

`getSpecies`という名前の関数を作成してください。

```js
assert.isFunction(getSpecies);
```

`getSpecies`関数は`animal`という単一のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getSpecies(tiger)`をコンソールにログ出力してください。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies`はドット記法を使ってspeciesプロパティにアクセスするべきです。

```js
assert.match(code, /animal\.species/);
```

`getSpecies(tiger)`を呼び出すと`"Tiger"`が返されるべきです。

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies`は関数のパラメータを使い、どんなオブジェクトでも扱えるようにしてください。

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
