---
id: 6999faf886634b72be7a819b
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

この野生動物トラッカーでは、新しいプロパティを使ってオブジェクトを更新します。

プロパティを持つオブジェクトを作成する例を示します。

```js
const animal = {
  name: "Lion" // property: value
};
```

次に、`tiger` オブジェクトを更新して、`species` というプロパティを含めてください。

その値を `"Tiger"` に設定してください。

# --hints--

`tiger` はオブジェクトである必要があります。

```js
assert.isObject(tiger);
```

`tiger` は `species` プロパティを持つ必要があります。

```js
assert.property(tiger, "species");
```

`tiger.species` は `"Tiger"` と等しくなる必要があります。

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
