---
id: 6999adc72c02812a28c01e31
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

このステップでは、オブジェクトのプロパティを更新する関数を作成します。

関数内でプロパティを更新する例を示します：

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

この例では、`age` プロパティが新しい値に更新されています。

次に、`updateAge` という名前の関数を作成してください。この関数は2つのパラメータ、`animal` と `newAge` を受け取ります。

関数内で、`age` オブジェクトの `animal` プロパティを `newAge` に更新してください。更新した `animal` オブジェクトを返します。

関数を作成したら、`console.log` を使って `updateAge(elephant, 12)` を呼び出し、更新された `elephant` オブジェクトをコンソールで確認してください。

# --hints--

`updateAge` という名前の関数を作成してください。

```js
assert.isFunction(updateAge);
```

`updateAge` 関数は2つのパラメータ、`animal` と `newAge` を持つ必要があります。

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` はドット記法を使って `age` プロパティを更新してください。

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

`updateAge` 関数は更新された `animal` オブジェクトを返す必要があります。

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

`updateAge(elephant, 12)` をコンソールにログ出力してください。

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

`updateAge(elephant, 12)` を呼び出すと age プロパティが更新されます。

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` は関数のパラメータを使い、どんなオブジェクトでも動作するようにしてください。

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
