---
id: 6999a0c1698b93d5d41156ed
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

`tiger` オブジェクトに `isEndangered` という名前の別のプロパティを追加してください。

`isEndangered` プロパティを `true` に設定してください。

# --hints--

`tiger` は `isEndangered` プロパティを持つべきです。

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` は `true` と等しくなるべきです。

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
