---
id: 699a00de9f564bb0effaa14d
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

このステップでは、オブジェクトの`age`プロパティにアクセスする関数を作成します。

例を示します：

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

この例では、関数がオブジェクトをパラメータとして受け取り、その`age`プロパティを返します。

では、`getAge`という名前の関数を作成してください。

関数は`animal`という名前のパラメータを1つ受け取ります。

ドット記法を使って、`age`の`animal`プロパティを返してください。

関数を作成したら、`console.log`を使って`getAge(tiger)`を呼び出し、コンソールで結果を確認してください。

# --hints--

`getAge`という名前の関数を作成してください。

```js
assert.isFunction(getAge);
```

`getAge`関数は、`animal`という単一のパラメータを持つべきです。

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge`はドット記法を使って`age`プロパティにアクセスするべきです。

```js
assert.match(code, /animal\.age/);
```

`getAge(tiger)`をコンソールにログ出力してください。

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

`getAge(tiger)`を呼び出すと`5`が返されるべきです。

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge`は関数のパラメータを使い、どんなオブジェクトでも動作するべきです。

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
