---
id: 69ba8da6b296391d913096a0
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

今度は `elephant` という名前の2つ目のオブジェクトを作成してください。

次のプロパティを追加してください：

- `species` に `"Elephant"` の値を設定する
- `age` に `10` の値を設定する
- `isEndangered` に `true` の値を設定する

# --hints--

`elephant` という名前の変数を作成してください。

```js
assert.exists(elephant);
```

`elephant` は `species` プロパティを持ち、その値は `"Elephant"` と等しくなければなりません。

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` は `age` プロパティを持ち、その値は `10` と等しくなければなりません。

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` は `isEndangered` プロパティを持ち、その値は `true` と等しくなければなりません。

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
