---
id: 683ecb010f4c6bafd428802c
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, convaliderai un insieme di dati medici per assicurarti che rispetti un insieme di regole.

Dichiara una variabile chiamata `medical_records` e assegnale una lista vuota. Nei passaggi successivi, la userai per archiviare i tuoi dati medici.

# --hints--

Dovresti dichiarare una variabile chiamata `medical_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

Dovresti assegnare una lista vuota alla tua variabile `medical_records`.

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
