---
id: 6842aee5bcaf4ab8883f8191
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Come hai imparato nelle lezioni precedenti, un dizionario è una struttura dati che contiene coppie chiave-valore, dove le chiavi sono tipi di dati immutabili (come le stringhe) e uniche all'interno di un dizionario:

```py
person = {
    'name': 'John',
    'age': 33
}
```

La lista `medical_records` conterrà dizionari, ognuno dei quali rappresenta un paziente. Aggiungi un dizionario con una chiave `patient_id` e un valore di tipo stringa `P1001` alla lista `medical_records`.

# --hints--

La tua lista `medical_records` dovrebbe contenere un solo dizionario.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Dovresti avere una chiave `patient_id` con il valore della stringa `P1001` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
