---
id: 6842ba0086201089ea268c92
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Seguendo la stessa struttura usata nei passaggi precedenti, la lista `medical_records` è stata compilata per te con i dati di altri pazienti. Sentiti libero di darci un’occhiata.

Ora inizierai a scrivere la funzione per validare il set di dati. Crea una funzione chiamata `validate` con un singolo parametro `data`.

Vuoi assicurarti che i tuoi dati siano una lista o una tupla. Quindi, all’interno della funzione `validate`, dichiara una variabile chiamata `is_sequence` e assegnale una chiamata a `isinstance`. Passa `data` come primo argomento e una tupla contenente `list` e `tuple` come secondo argomento.

# --hints--

Dovresti avere una funzione chiamata `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

La tua funzione `validate` dovrebbe avere un singolo parametro `data`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

Dovresti avere una variabile chiamata `is_sequence` all’interno della tua funzione `validate`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

Dovresti assegnare `isinstance(data, (list, tuple))` alla tua variabile `is_sequence`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
