---
id: 6842d5b537c64266420faf62
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Aggiungi una chiave `age` con il valore dell'intero `34` e una chiave `gender` con il valore della stringa `Female` al tuo dizionario. Non dimenticare la virgola tra le coppie chiave-valore.

# --hints--

La tua lista `medical_records` dovrebbe contenere un singolo dizionario.

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

Dovresti avere una chiave `age` con il valore dell'intero `34` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Dovresti avere una chiave `gender` con il valore della stringa `Female` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

Dovresti avere una chiave `patient_id` con il valore della stringa `P1001` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
