---
id: 6842d659a4ae7973dabab3b0
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Completa il dizionario aggiungendo le seguenti tre coppie chiave-valore:

- Una chiave `diagnosis` con il valore `Hypertension`
- Una chiave `medications` con il valore della lista `['Lisinopril']`
- Una chiave `last_visit_id` con il valore `V2301`.

# --hints--

Dovresti avere una chiave `diagnosis` con il valore della stringa `Hypertension` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

Dovresti avere una chiave `medications` con il valore della lista `['Lisinopril']` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

Dovresti avere una chiave `last_visit_id` con il valore della stringa `V2301` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

Dovresti avere una chiave `patient_id` con il valore della stringa `P1001` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

Dovresti avere una chiave `age` con il valore dell'intero `34` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

Dovresti avere una chiave `gender` con il valore della stringa `Female` nel tuo dizionario.

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
