---
id: 6846b0048f71fc4efd786b42
title: Passaggio 18
challengeType: 20
dashedName: step-18
---

# --description--

Il metodo `keys()` restituisce un oggetto view contenente tutte le chiavi di un dizionario:

```py
person = {
   'name': 'John',
   'age': 33
}

print(person.keys()) # dict_keys(['name, 'age'])
```

All'interno del tuo loop `for`, dopo la prima istruzione `if`, crea un'istruzione `if` che venga eseguita quando l'insieme delle chiavi del dizionario corrente è diverso da `key_set`. Questo serve a garantire che non ci siano chiavi mancanti o non valide nel dizionario.

All'interno della nuova istruzione `if`, stampa `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.` (dove `<dictionary>` e `<index>` devono essere sostituiti dal dizionario e dall'indice nell'iterazione corrente) e imposta `is_invalid` su `True`.

# --hints--

Dovresti avere una seconda istruzione `if` all'interno del tuo loop `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1]
`) })
```

La tua nuova istruzione `if` dovrebbe essere eseguita quando `set(dictionary.keys())` è diverso da `key_set`.

```js
({ test: () => assert(runPython(`
_condition = _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_conditions()[0]

_condition.is_equivalent("set(dictionary.keys()) != key_set") or _condition.is_equivalent("key_set != set(dictionary.keys())")
`)) })
```

Dovresti stampare `Invalid format: <dictionary> at position <index> has missing and/or invalid keys.` (dove `<dictionary>` e `<index>` devono essere sostituiti dal dizionario e dall'indice nell'iterazione corrente) all'interno della tua nuova istruzione `if`.

```js
({ test: () => assert(runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
_test_data = [
  {'testing': 'stuff'},
  {
    'patient_id': 'P1003',
    'age': 29,
    'gender': 'Female',
    'diagnosis': 'Asthma',
    'medications': ['Albuterol'],
    'last_visit_id': 'V2303',
  },
  {'keep testing': 'stuff'}
]
validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
len(out) >= 2 and _first in out and _second in out
`)) })
```

Dovresti impostare `is_invalid` su `True` all'interno della tua nuova istruzione `if`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_ifs()[1].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }   
]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

--fcc-editable-region--
        
--fcc-editable-region--

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
