---
id: 6846e3e61271e1f164636b13
title: Passaggio 21
challengeType: 20
dashedName: step-21
---

# --description--

Ora renderai la validazione più dettagliata. Crea una funzione chiamata `find_invalid_records` per trovare valori non validi in un dizionario. Dagli i seguenti parametri: `patient_id`, `age`, `gender`, `diagnosis`, `medications`, `last_visit_id`.

All'interno della tua nuova funzione, crea un dizionario vuoto chiamato `constraints`. Poi, restituisci `constraints` dalla tua nuova funzione.

# --hints--

Dovresti avere una funzione chiamata `find_invalid_records`.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("find_invalid_records")
`) })
```

La tua funzione `find_invalid_records` dovrebbe avere i seguenti parametri: `patient_id`, `age`, `gender`, `diagnosis`, `medications`, `last_visit_id`, in questo ordine.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").has_args("patient_id, age, gender, diagnosis, medications, last_visit_id")
`) })
```

Dovresti avere una variabile chiamata `constraints` con il valore di un dizionario vuoto all'interno della tua nuova funzione.

```js
({ test: () => assert(runPython(`
_foo = _Node(_code).find_function("find_invalid_records")
_foo.has_stmt("constraints = {}") or _foo.has_stmt("constraints = dict()")
`)) })
```

Dovresti restituire il dizionario `constraints` dalla tua funzione `find_invalid_records`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").has_return("constraints")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
