---
id: 684ed67f4c678896445f973f
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

L'operatore `**` può essere usato per decomprimere gli elementi in un dizionario e passarli come argomenti con nome in una chiamata di funzione:

```py
def sum(a, b, c):
    return a + b + c
nums = {'a': 2, 'b': 4, 'c': 1}

print(sum(**nums)) # 7
```

Nell'esempio sopra, `sum(**nums)` è equivalente a `sum(a=2, b=4, c=1)`.

Alla fine del tuo codice, stampa il risultato della chiamata alla funzione `find_invalid_records`. Per i suoi argomenti, usa l'operatore `**` per decomprimere `medical_records[0]`.

# --hints--

Dovresti stampare `find_invalid_records(**medical_records[0])`.

```js
({ test: () => runPython(`
assert _Node(_code).has_call("print(find_invalid_records(**medical_records[0]))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
--fcc-editable-region--

--fcc-editable-region--
```
