---
id: 68529fb26c581f13e901de3a
title: Passaggio 44
challengeType: 20
dashedName: step-44
---

# --description--

Subito dopo la variabile `invalid_records`, crea un loop `for` per iterare su di essa. Per ogni record non valido, stampa `Unexpected format '<key>: <val>' at position <index>.`. Sostituisci `<key>`, `<val>` e `<index>` con la chiave, il valore e l’indice correnti.

Ricorda che `invalid_records` è una lista di chiavi che si riferiscono ai record non validi nell’attuale `dictionary`. Dovrai prendere la chiave da `invalid_records` e cercare il valore in `dictionary`.

Position o `index` indica il dizionario corrente in `medical_records`, definito dal loop esterno `for` nella funzione.

Rivedi il tuo codice finora se ti serve ricordare i loop e le variabili già create.

Poi, imposta `is_invalid` su `True`.

Sentiti libero di testare la funzione `validate` con dati non validi per vedere i messaggi di validazione.

Con questo, il workshop del validatore medico è completo.

# --hints--

Dovresti avere un loop `for` che itera su `invalid_records`.

```js
({ test: () => assert(runPython(`
_Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_for_iter().is_equivalent("invalid_records")
`)) })
```

Dovresti stampare `Unexpected format '<key>: <val>' at position <index>.` (dove `<key>`, `<val>` e `<index>` devono essere sostituiti dalla chiave, valore e indice correnti) all’interno del tuo loop `for`.

```js
({ test: () => runPython(`
built_in_print = print
out = []

def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
_test_data = [
  {'testing': 'stuff'},
  {
    'patient_id': 'P1003',
    'age': 29,
    'gender': 'emale',
    'diagnosis': 'Asthma',
    'medications': ['Albuterol'],
    'last_visit_id': 'V2303',
  },
  {'keep testing': 'stuff'},
  {
    'patient_id': '1003',
    'age': 29,
    'gender': 'female',
    'diagnosis': 'Asthma',
    'medications': 'Albuterol',
    'last_visit_id': 'V2303',
  },
]

validate(_test_data)
print = built_in_print

_first = f"Invalid format: {_test_data[0]} at position 0 has missing and/or invalid keys."
_second = "Unexpected format 'gender: emale' at position 1."
_third = f"Invalid format: {_test_data[2]} at position 2 has missing and/or invalid keys."
_fourth = "Unexpected format 'patient_id: 1003' at position 3."
_fifth = "Unexpected format 'medications: Albuterol' at position 3."

assert len(out) >= 5
assert _first in out
assert _second in out
assert _third in out
assert _fourth in out
assert _fifth in out
`) })
```

Dovresti impostare `is_invalid` su `True` all’interno del tuo nuovo loop `for`.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].has_stmt("is_invalid = True")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
    return [key for key, value in constraints.items() if not value]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True
            continue

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True
            continue

        invalid_records = find_invalid_records(**dictionary)
--fcc-editable-region--
        
--fcc-editable-region--

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```

# --solutions--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
        'patient_id': isinstance(patient_id, str)
        and re.fullmatch('p\d+', patient_id, re.IGNORECASE),
        'age': isinstance(age, int) and age >= 18,
        'gender': isinstance(gender, str) and gender.lower() in ('male', 'female'),
        'diagnosis': isinstance(diagnosis, str) or diagnosis is None,
        'medications': isinstance(medications, list)
        and all([isinstance(i, str) for i in medications]),
        'last_visit_id': isinstance(last_visit_id, str)
        and re.fullmatch('v\d+', last_visit_id, re.IGNORECASE)
    }
    return [key for key, value in constraints.items() if not value]

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True
            continue

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True
            continue

        invalid_records = find_invalid_records(**dictionary)
        for key in invalid_records:
            print(
                f"Unexpected format '{key}: {dictionary.get(key)}' at position {index}."
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
```
