---
id: 56533eb9ac21ba0edf2244e2
title: Caesars Chiffre
challengeType: 5
saveSubmissionToDB: true
forumTopicId: 16003
dashedName: caesars-cipher
---

# --description--

Eine der einfachsten und bekanntesten <dfn>Chiffren</dfn> ist die <dfn>Caesar-Chiffre</dfn>, auch bekannt als <dfn>Verschiebungs-Chiffre</dfn>. Bei einer Verschiebungs-Chiffre werden die Bedeutungen der Buchstaben um einen festen Wert verschoben.

Eine gängige moderne Anwendung ist die <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a>-Chiffre, bei der die Buchstabenwerte um 13 Stellen verschoben werden. Somit `A ↔ N`, `B ↔ O` und so weiter.

Schreiben Sie eine Funktion, die einen mit <a href="https://www.freecodecamp.org/news/how-to-code-the-caesar-cipher-an-introduction-to-basic-encryption-3bf77b4e19f7/" target="_blank" rel="noopener noreferrer nofollow">ROT13</a> codierten String als Eingabe erhält und einen decodierten String zurückgibt.

Alle Buchstaben sind großgeschrieben. Verändern Sie keine nicht-alfabetischen Zeichen (z. B. Leerzeichen, Satzzeichen), sondern geben Sie diese unverändert weiter.

# --hints--

`rot13("SERR PBQR PNZC")` sollte in den String `FREE CODE CAMP` decodiert werden

```js
assert(rot13('SERR PBQR PNZC') === 'FREE CODE CAMP');
```

`rot13("SERR CVMMN!")` sollte in den String `FREE PIZZA!` decodiert werden

```js
assert(rot13('SERR CVMMN!') === 'FREE PIZZA!');
```

`rot13("SERR YBIR?")` sollte in den String `FREE LOVE?` decodiert werden

```js
assert(rot13('SERR YBIR?') === 'FREE LOVE?');
```

`rot13("GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.")` sollte in den String `THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.` decodiert werden

```js
assert(
  rot13('GUR DHVPX OEBJA SBK WHZCF BIRE GUR YNML QBT.') ===
    'THE QUICK BROWN FOX JUMPS OVER THE LAZY DOG.'
);
```

# --seed--

## --seed-contents--

```js
function rot13(str) {
  return str;
}

rot13("SERR PBQR PNZC");
```

# --solutions--

```js
var lookup = {
  'A': 'N','B': 'O','C': 'P','D': 'Q',
  'E': 'R','F': 'S','G': 'T','H': 'U',
  'I': 'V','J': 'W','K': 'X','L': 'Y',
  'M': 'Z','N': 'A','O': 'B','P': 'C',
  'Q': 'D','R': 'E','S': 'F','T': 'G',
  'U': 'H','V': 'I','W': 'J','X': 'K',
  'Y': 'L','Z': 'M'
};

function rot13(encodedStr) {
  var codeArr = encodedStr.split("");  // String to Array
  var decodedArr = []; // Your Result goes here
  // Only change code below this line

  decodedArr = codeArr.map(function(letter) {
    if(lookup.hasOwnProperty(letter)) {
      letter = lookup[letter];
    }
    return letter;
  });

  // Only change code above this line
  return decodedArr.join(""); // Array to String
}
```
