---
id: 69447054012a22f38c944ff0
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

在你的 `countup` 函数内添加一个 `if` 语句。

查看 `number` 是否小于 `1` 并返回一个空数组。

这将作为你递归的基本情形。

# --hints--

当 `number` 小于 `1` 时，函数应返回一个空数组。

```js
assert.deepStrictEqual(countup(0), []);
```

当 `number` 不小于 `1` 时，函数不应返回任何内容。

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
