---
id: 694470b9a65b278f981ec62e
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

最后，通过在 `console.log()` 语句中使用参数 `5` 调用你的 `countup` 函数来测试它。 

输出应为 `[1, 2, 3, 4, 5]`。

# --hints--

你应该在 `console.log()` 语句中用参数 `5` 调用函数 `countup`。

```js
assert.deepStrictEqual(countup(5), [1, 2, 3, 4, 5]);
assert.match(code, /console\.log\s*\(\s*countup\s*\(\s*5\s*\)\s*\)/)
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
  } else {
    countArray = countup(number - 1);
    countArray.push(number);
    return countArray;
  }
}
console.log(countup(5));
```
