---
id: 69446f18fdcb11aa516f7c5c
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将创建一个 `countup` 函数，该函数使用递归创建一个包含从 1 到给定数字的数字的数组。

首先创建一个名为 `countup` 的空函数，该函数接受一个参数。

# --hints--

你应该有一个名为 `countup` 的函数。

```js
assert.isFunction(countup);
```

你的 `countup` 函数应接受一个参数。

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
