---
id: 694870bcf3c6874112b6f951
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

对于所有大于且等于 `1` 的数字，创建一个 `else` 块，并在其中调用 `countup` 函数，参数为 `number - 1`，并将结果赋值给 `countArray`。 

然后使用 `.push()` 将 `number` 添加到 `countArray`，最后返回 `countArray`。 

这确保递归最终达到基本情形并终止。

# --hints--

你应该将调用 `countup` 并传入 `number - 1` 的结果赋值给 `countArray`。

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

你应该将 `number` 推入 `countArray`。

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

你应该返回 `countArray`。

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
