---
id: 6948705160b66cf93a916e29
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在 `countup` 函数内，在你的 `if` 语句正上方声明一个名为 `countArray` 的空数组。

这个数组将在每次函数调用时保存 `number` 的值。

# --hints--

你应该声明一个名为 `countArray` 的空数组。

```js
const explorer = await __helpers.Explorer(code);
const { countup } = explorer.functions;
assert.exists(countup);
const { countArray } = countup.variables;
assert.exists(countArray);
assert.isTrue(countArray.value.matches('[]'));
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
