---
id: 5d5a813321b9e3db6c106a46
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

JavaScript è un linguaggio potente che ti permette di creare siti web interattivi.

*Nota*: Per tutti i progetti rimanenti in questo curriculum, ti servirà una conoscenza di base di HTML e CSS. Se sei nuovo a HTML e CSS, segui il [Responsive Web Design Certification](https://www.freecodecamp.org/learn/2022/responsive-web-design/).

Per iniziare, crea il tuo codice boilerplate HTML standard con un `DOCTYPE`, `html`, `head` e `body`, poi aggiungi un tag `meta` per il `charset`. Aggiungi un elemento `title` e usa il testo `RPG - Dragon Repeller` per esso. Includi un tag `link` per il tuo foglio di stile per collegare il file `styles.css`.

Infine, crea un elemento `div` con `id` impostato su `game` all’interno del tuo `body`.

# --hints--

Il tuo codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un `>` dopo il tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di apertura. Non dimenticare l’attributo `lang`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all’inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `head`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag di apertura `body`.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert.match(code, /<head\s*>.*<\/head\s*>.*<body\s*>.*<\/body\s*>/s)
```

L’elemento `head` dovrebbe essere all’interno dell’elemento `html`.

```js
assert.match(code, /<html[^>]*>.*<head\s*>.*<\/head\s*>.*<\/html\s*>/s);
```

L’elemento `body` dovrebbe essere all’interno dell’elemento `html`.

```js
assert.match(code, /<html[^>]*>.*<body\s*>.*<\/body\s*>.*<\/html\s*>/s);
```

Il tuo codice dovrebbe avere un elemento `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Il tuo elemento `meta` dovrebbe avere un attributo `charset` con valore `UTF-8`.

```js
assert.match(code, /<meta[\s\S]+?charset=('|"|`)UTF-8\1/i)
```

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Dovresti avere un tag di chiusura `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Il tuo codice dovrebbe avere un elemento `link`.

```js
const link = document.querySelector('head > link');
assert.exists(link);
```

Dovresti avere un elemento `div`.

```js
const div = document.querySelector('div');
assert.exists(div);
```

Il tuo elemento `div` dovrebbe avere un attributo `id` con valore `game`.

```js
const div = document.querySelector('div');
assert.equal(div?.id, 'game');
```

Il tuo elemento `div` dovrebbe essere all’interno dell’elemento `body`.

```js
const div = document.querySelector('div');
assert.equal(div?.parentElement?.localName, 'body');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
