---
id: 62a1166ed9a56d439c0770e7
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Crea tre elementi `span` all’interno del tuo elemento `#stats`. Dai a ciascuno la classe `stat`. Poi assegna al primo il testo `XP: 0`, al secondo il testo `Health: 100` e al terzo il testo `Gold: 50`.

# --hints--

Dovresti avere tre elementi `span` all’interno dell’elemento `#stats`.

```js
const spans = document.querySelectorAll('#stats > span');
assert.equal(spans?.length, 3);
```

Dai ai nuovi tre elementi `span` la classe `stat`.

```js
assert.lengthOf(document.querySelectorAll('#stats > .stat'), 3);
```

Il tuo primo elemento `.stat` dovrebbe avere il testo fornito `XP: 0`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[0]?.innerText === 'XP: 0');
```

Il tuo secondo elemento `.stat` dovrebbe avere il testo fornito `Health: 100`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[1]?.innerText === 'Health: 100');
```

Il tuo terzo elemento `.stat` dovrebbe avere il testo fornito `Gold: 50`.

```js
assert(document.querySelectorAll('#stats > .stat')?.[2]?.innerText === 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
      <div id="stats">

      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
