---
id: 62a115879a6d51422652cbfc
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Nel tuo gioco di ruolo, gli utenti potranno tenere traccia delle loro statistiche, comprare armi e combattere mostri. Prima di poter continuare con la parte interattiva in JavaScript del gioco, devi prima creare gli elementi HTML che mostreranno le informazioni del gioco.

Crea quattro elementi `div` all'interno del tuo elemento `#game`. Assegna loro i seguenti rispettivi valori `id`, in ordine: `stats`, `controls`, `monsterStats` e `text`.

# --hints--

Dovresti creare quattro nuovi elementi `div`.

```js
assert.equal(document.querySelectorAll('div')?.length, 5);
```

Dovresti dare a uno dei nuovi elementi `div` un `id` di `stats`.

```js
assert.exists(document.querySelector('div#stats'));
```

Dovresti dare a uno dei nuovi elementi `div` un `id` di `controls`.

```js
assert.exists(document.querySelector('div#controls'));
```

Dovresti dare a uno dei nuovi elementi `div` un `id` di `monsterStats`.

```js
assert.exists(document.querySelector('div#monsterStats'));
```

Dovresti dare a uno dei nuovi elementi `div` un `id` di `text`.

```js
assert.exists(document.querySelector('div#text'));
```

Dovresti posizionare i nuovi elementi `div` nell'ordine corretto.

```js
function __t(a, b) {
  return document.querySelector(a)?.nextElementSibling?.getAttribute('id') === b;
}
assert(__t('div#stats','controls') && __t('div#controls','monsterStats') && __t('div#monsterStats','text'));
```

Dovresti inserire i nuovi elementi `div` all'interno dell'elemento `#game`.

```js
assert.equal(document.querySelector('#game')?.children?.length, 4);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
--fcc-editable-region--
  <body>
    <div id="game">
    </div>
  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
