---
id: 62a2509ba163e020bb9d84ea
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Ora puoi iniziare a scrivere il tuo JavaScript. Inizia creando un elemento `script`. Questo elemento serve per caricare JavaScript nel tuo file HTML.

```html
<script>
  // JavaScript code goes here
</script>
```

# --hints--

Dovresti avere un elemento `script`.

```js
assert.isAtLeast(document.querySelectorAll('script').length, 1);
```

Il tuo elemento `script` dovrebbe avere un tag di apertura.

```js
assert.match(code, /<script\s*>/i);
```

Il tuo elemento `script` dovrebbe avere un tag di chiusura.

```js
assert.match(code, /<\/script\s*>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
