---
id: 62a23c1d505bfa13747c8a9b
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Avvolgi i numeri `0`, `100` e `50` in elementi `span`, e avvolgi questi nuovi elementi `span` in elementi `strong`. Poi assegna ai tuoi nuovi elementi `span` valori `id` di `xpText`, `healthText` e `goldText`, rispettivamente.

La tua risposta dovrebbe seguire questa struttura di base:

```html
<span class="stat">TEXT <strong><span id="VALUE">TEXT</span></strong></span>
```

# --hints--

Dovresti aggiungere un elemento `strong` nel tuo primo elemento `.stat`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Il tuo primo nuovo elemento `strong` dovrebbe avere un elemento `span`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Il tuo primo elemento `span` annidato dovrebbe avere il `id` di `xpText`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'xpText');
```

Il tuo primo elemento `span` dovrebbe avvolgere il testo `0`.

```js
const stat = document.querySelectorAll('.stat')[0];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '0');
```

Il testo del tuo primo elemento `.stat` dovrebbe essere ancora `XP: 0`.

```js
const stat = document.querySelectorAll('.stat')[0];
assert.equal(stat.innerText.trimEnd(), 'XP: 0');
```

Dovresti aggiungere un elemento `strong` nel tuo secondo elemento `.stat`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Il tuo secondo nuovo elemento `strong` dovrebbe avere un elemento `span`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Il tuo secondo elemento `span` annidato dovrebbe avere il `id` di `healthText`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'healthText');
```

Il tuo secondo elemento `span` dovrebbe avvolgere il testo `100`.

```js
const stat = document.querySelectorAll('.stat')[1];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '100');
```

Il testo del tuo secondo elemento `.stat` dovrebbe essere ancora `Health: 100`.

```js
const stat = document.querySelectorAll('.stat')[1];
assert.equal(stat.innerText.trimEnd(), 'Health: 100');
```

Dovresti aggiungere un elemento `strong` nel tuo terzo elemento `.stat`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
assert.exists(strong);
```

Il tuo terzo nuovo elemento `strong` dovrebbe avere un elemento `span`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(span);
```

Il tuo terzo elemento `span` annidato dovrebbe avere il `id` di `goldText`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span?.id, 'goldText');
```

Il tuo terzo elemento `span` dovrebbe avvolgere il testo `50`.

```js
const stat = document.querySelectorAll('.stat')[2];
const strong = stat?.querySelector('strong');
const span = strong?.querySelector('span');
assert.equal(span.innerText.trim(), '50');
```

Il testo del tuo terzo elemento `.stat` dovrebbe essere ancora `Gold: 50`.

```js
const stat = document.querySelectorAll('.stat')[2];
assert.equal(stat.innerText.trimEnd(), 'Gold: 50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
--fcc-editable-region--
        <span class="stat">XP: 0</span>
        <span class="stat">Health: 100</span>
        <span class="stat">Gold: 50</span>
--fcc-editable-region--
      </div>
      <div id="controls"></div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
