---
id: 62a240c67f3dbb1a1e6d95ee
title: Passaggio 23
challengeType: 0
dashedName: step-23
---

# --description--

Dai al tuo `#game` una larghezza massima di `500px` e un'altezza massima di `400px`. Imposta il `background-color` su `#ffffff` e il `color` su `#ffffff`. 

Usa i margini per centrarlo impostando il margine superiore su `30px`, il margine inferiore su `0px` e i margini sinistro e destro su `auto`. 

Infine, dagli `10px` di spaziatura interna su tutti e quattro i lati.

# --hints--

Dovresti avere un selettore `#game`.

```js
const game = new __helpers.CSSHelp(document).getStyle('#game');
assert.exists(game);
```

Il tuo selettore `#game` dovrebbe avere un `max-width` di `500px`.

```js
const maxWidth = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-width');
assert.equal(maxWidth, '500px');
```

Il tuo selettore `#game` dovrebbe avere un `max-height` di `400px`.

```js
const maxHeight = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('max-height');
assert.equal(maxHeight, '400px');
```

Il tuo selettore `#game` dovrebbe avere un `background-color` di `#ffffff`.

```js
const background = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('background-color');
assert.equal(background, 'rgb(255, 255, 255)');
```

Il tuo selettore `#game` dovrebbe avere un `color` di `#ffffff`.

```js
const color = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('color');
assert.equal(color, 'rgb(255, 255, 255)');
```

Il tuo selettore `#game` dovrebbe avere un `margin-top` impostato su `30px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-top');
assert.equal(margin, '30px');
```

Il tuo selettore `#game` dovrebbe avere un `margin-left` impostato su `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-left');
assert.equal(margin, 'auto');
```

Il tuo selettore `#game` dovrebbe avere un `margin-right` impostato su `auto`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-right');
assert.equal(margin, 'auto');
```


Il tuo selettore `#game` dovrebbe avere un `margin-bottom` impostato su `0px`.

```js
const margin = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('margin-bottom');
assert.equal(margin, '0px');
```

Il tuo selettore `#game` dovrebbe avere `10px` di spaziatura interna su tutti i lati.

```js
const padding = new __helpers.CSSHelp(document).getStyle('#game')?.getPropertyValue('padding');
assert.equal(padding, '10px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats">
        <span class="stat">Monster Name: <strong><span id="monsterName"></span></strong></span>
        <span class="stat">Health: <strong><span id="monsterHealth"></span></strong></span>
      </div>
      <div id="text">
        Welcome to Dragon Repeller. You must defeat the dragon that is preventing people from leaving the town. You are in the town square. Where do you want to go? Use the buttons above.
      </div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
body {
  background-color: #0a0a23;
}

#text {
  background-color: #0a0a23;
  color: #ffffff;
  padding: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
