---
id: 62a23cb9bc467a147516b500
title: Passaggio 14
challengeType: 0
dashedName: step-14
---

# --description--

Per il tuo elemento `#controls`, crea tre elementi `button`. Il primo dovrebbe avere la proprietà `id` impostata su `button1`, e il testo `Go to store`. Il secondo dovrebbe avere la proprietà `id` impostata su `button2`, e il testo `Go to cave`. Il terzo dovrebbe avere la proprietà `id` impostata su `button3`, e il testo `Fight dragon`.

# --hints--

Dovresti aggiungere tre elementi `button` al tuo elemento `#controls`.

```js
const buttons = document.querySelectorAll('#controls > button');
assert.exists(buttons);
assert.equal(buttons.length, 3);
```

Il tuo primo pulsante dovrebbe avere la proprietà `id` impostata su `button1`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.id, 'button1');
```

Il tuo primo pulsante dovrebbe avere il testo `Go to store`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button1 = buttons[0];
assert.equal(button1.innerText, 'Go to store');
```

Il tuo secondo pulsante dovrebbe avere la proprietà `id` impostata su `button2`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.id, 'button2');
```

Il tuo secondo pulsante dovrebbe avere il testo `Go to cave`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button2 = buttons[1];
assert.equal(button2.innerText, 'Go to cave');
```

Il tuo terzo pulsante dovrebbe avere la proprietà `id` impostata su `button3`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.id, 'button3');
```

Il tuo terzo pulsante dovrebbe avere il testo `Fight dragon`.

```js
const buttons = document.querySelectorAll('#controls > button');
const button3 = buttons[2];
assert.equal(button3.innerText, 'Fight dragon');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
--fcc-editable-region--
      <div id="controls">

      </div>
--fcc-editable-region--
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];
```
