---
id: 62a257659d0d1e2456f24ba2
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Prima di iniziare a scrivere il codice del progetto, spostalo in un file a parte per mantenere tutto organizzato.

Rimuovi la riga `console.log("Hello World");`. Poi assegna al tuo elemento `script` ora vuoto un attributo `src` impostato su `./script.js`.

# --hints--

Non dovresti avere una riga `console.log("Hello World");` nel tuo codice.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Il tuo elemento `script` dovrebbe avere un attributo `src` impostato su `./script.js`.

```js
const script = document.querySelector("script[data-src$='script.js']");
assert.exists(script);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>
      console.log("Hello World");
    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
