---
id: 62a23d1c5f1c93161f3582ae
title: Passaggio 19
challengeType: 0
dashedName: step-19
---

# --description--

Simile al tuo elemento `#stats`, il tuo elemento `#monsterStats` ha bisogno di due elementi `span`. Dagli la classe `stat` e assegna al primo elemento il testo `Monster Name: ` e al secondo il testo `Health: `. Dopo il testo di ciascuno, aggiungi un elemento `strong` con un elemento `span` annidato vuoto. Dai al primo elemento `span` interno un `id` di `monsterName` e al secondo elemento `span` interno un `id` di `monsterHealth`.

# --hints--

Il tuo elemento `#monsterStats` dovrebbe avere due elementi `span`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans.length, 2);
```

I tuoi nuovi elementi `span` dovrebbero entrambi avere un valore `class` di `stat`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.equal(spans[0].className, 'stat');
assert.equal(spans[1].className, 'stat');
```

Il tuo primo elemento `span` dovrebbe avere un elemento `strong` con un elemento `span` annidato vuoto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[0]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Il tuo secondo elemento `span` dovrebbe avere un elemento `strong` con un elemento `span` annidato vuoto.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const strong = spans[1]?.querySelector('strong');
const span = strong?.querySelector('span');
assert.exists(strong);
assert.exists(span);
```

Il tuo primo elemento `span` dovrebbe avere il testo `Monster Name: `. Assicurati che gli spazi siano corretti.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+/);
```

Il tuo secondo elemento `span` dovrebbe avere il testo `Health: `. Assicurati che gli spazi siano corretti.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[1].innerHTML, /Health:\s+/);
```

I tuoi elementi `strong` e `span` dovrebbero venire dopo il testo.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
assert.match(spans[0].innerHTML, /Monster Name:\s+<strong>/);
assert.match(spans[1].innerHTML, /Health:\s+<strong>/);
```

Il tuo primo elemento `span` interno dovrebbe avere un `id` di `monsterName`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[0]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterName');
```

Il tuo secondo elemento `span` interno dovrebbe avere un `id` di `monsterHealth`.

```js
const spans = document.querySelectorAll(`#monsterStats > span`);
const span = spans[1]?.querySelector('span');
const spanId = span?.id;
assert.equal(spanId, 'monsterHealth');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
--fcc-editable-region--
      <div id="monsterStats">

      </div>
--fcc-editable-region--
      <div id="text"></div>
    </div>
    <script src="./script.js"></script>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

const button1 = document.querySelector("#button1");
const button2 = document.querySelector("#button2");
const button3 = document.querySelector("#button3");
```
