---
id: 62a255dae245b52317da824a
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Uno degli strumenti più potenti è la console per sviluppatori. A seconda del browser, potresti aprirla premendo `F12` o `Ctrl+Shift+I`. Su Mac, puoi premere `Option + ⌘ + C` e selezionare "Console". Puoi anche cliccare il pulsante "Console" sopra la finestra di anteprima per vedere la nostra console integrata.

La console per sviluppatori mostrerà gli errori prodotti dal tuo codice, ma puoi anche usarla per vedere i valori delle variabili nel tuo codice, cosa utile per eseguire il debug.

Aggiungi una riga `console.log("Hello World");` tra i tuoi tag `script`. Poi clicca il pulsante "Console" per aprire la console. Dovresti vedere il testo `"Hello World"`.


# --hints--

Dovresti aggiungere una riga `console.log("Hello World");` al tuo codice.

```js
assert.match(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;?/);
```

La tua riga `console.log("Hello World");` dovrebbe essere tra i tuoi tag `script`.

```js
assert.match(code, /<script\s*>\s*console\.log\(\s*('|")Hello World\1\s*\)\s*;?\s*<\/script>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
--fcc-editable-region--
    <script>

    </script>
--fcc-editable-region--
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```
