---
id: 62a3a7e4f1060e2fc5ffb34b
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Crea un'altra variabile chiamata `currentWeaponIndex` e impostala su `0`.

# --hints--

Dovresti usare `let` per dichiarare una variabile chiamata `currentWeaponIndex`.

```js
assert.match(code, /let\s+currentWeaponIndex/i);
```

Dovresti usare camelCase per nominare la tua variabile.

```js
assert.match(code, /currentWeaponIndex/);
```

La tua variabile `currentWeaponIndex` dovrebbe essere impostata su `0`.

```js
assert.equal(currentWeaponIndex, 0);
```

Dovresti inizializzare la tua variabile a `0`.

```js
assert.match(code, /let\s+currentWeaponIndex\s*=\s*0/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
--fcc-editable-region--
```
