---
id: 62a3a0a3c0a4b32915d26a6e
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Rimuovi la riga `console.log("Hello World");` per iniziare a scrivere il codice del tuo progetto.

Nel progetto precedente, hai imparato a lavorare con le variabili e la parola chiave `let` in questo modo:

```js
let camperbot;
```

Hai anche imparato come inizializzare una variabile con un valore in questo modo:

```js
let age = 32;
```

Usa la parola chiave `let` per dichiarare una variabile chiamata `xp` e assegnarle il valore `0`, un numero.

# --hints--

Non dovresti avere una riga `console.log("Hello World");` nel tuo codice.

```js
assert.notMatch(code, /console\.log\(\s*('|")Hello World\1\s*\)\s*;/);
```

Dovresti usare la parola chiave `let` per dichiarare la tua variabile.

```js
assert.match(code, /let/);
```

Dovresti dichiarare una variabile chiamata `xp`.

```js
assert.match(code, /let\s+xp/);
```

`xp` dovrebbe avere un valore di `0`.

```js
assert.equal(xp, 0);
```

Dovresti inizializzare la variabile `xp` a `0`.

```js
assert.match(code, /let\s+xp\s*=\s*0\s*;?/)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
console.log("Hello World");
--fcc-editable-region--
```
