---
id: 62a3b79d520a7f3d0e25afd6
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

JavaScript interagisce con l’HTML usando il <dfn>Document Object Model</dfn>, o DOM. Il DOM è un albero di oggetti che rappresenta l’HTML. Puoi consultare l’HTML usando l’oggetto `document`, che rappresenta l’intero documento HTML.

Un metodo per trovare elementi specifici nel tuo HTML è usare il metodo `querySelector()`. Il metodo `querySelector()` prende come argomento un selettore CSS e restituisce il primo elemento che corrisponde a quel selettore. Per esempio, per trovare l’elemento `<h1>` nel tuo HTML, scriveresti:

```js
let h1 = document.querySelector("h1");
```

Nota che `h1` è una stringa e corrisponde al selettore CSS che useresti.

Crea una variabile `button1` e usa `querySelector()` per assegnarle il tuo elemento con l’id `id` di `button1`. Ricorda che i selettori CSS `id` sono preceduti da un `#`.

# --hints--

Dovresti usare `let` per dichiarare una variabile `button1`.

```js
assert.match(code, /let\s+button1/);
```

Dovresti usare `document.querySelector()`.

```js
assert.match(code, /document\.querySelector/);
```

Dovresti usare il selettore `"#button1"`.

```js
assert.match(code, /querySelector\(\s*('|")#button1\1\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>
  </body>
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

--fcc-editable-region--

--fcc-editable-region--
```
