---
id: 62a3b41c9494f937560640ab
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Dichiara altre due variabili chiamate `monsterHealth` e `inventory`.

Per la tua variabile `inventory`, assegnale il valore di un array contenente la stringa `"stick"`.

Ricorda che nel progetto precedente hai lavorato con gli array in questo modo:

```js
let exampleArray = ["first", "second", "third"];
```

# --hints--

Dovresti usare `let` per dichiarare la tua variabile `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth/i);
```

Dovresti usare camelCase per nominare la tua variabile `monsterHealth`.

```js
assert.match(code, /monsterHealth/);
```

`monsterHealth` non dovrebbe avere un valore.

```js
assert.isUndefined(monsterHealth);
```

Non dovresti assegnare un valore alla tua variabile `monsterHealth`.

```js
assert.match(code, /let\s+monsterHealth\s*;?/);
```

Dovresti usare `let` per dichiarare la tua variabile `inventory`.

```js
assert.match(code, /let\s+inventory/i);
```

La tua variabile `inventory` dovrebbe essere ancora un array.

```js
assert.isArray(inventory);
```

La tua variabile `inventory` dovrebbe avere solo un valore.

```js
assert.lengthOf(inventory, 1);
```

La tua variabile `inventory` dovrebbe includere la stringa `"stick"`.

```js
assert.include(inventory, "stick");
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
    </div>
  </body>
</html>
```

```js
--fcc-editable-region--
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
--fcc-editable-region--
```
