---
id: 62b46e3a8d4be31be5af793d
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Abbiamo incontrato un piccolo problema. Stai cercando di fare una query nella tua pagina per un elemento button, ma il tuo tag `script` si trova nel `head` del tuo HTML. Questo significa che il tuo codice viene eseguito prima che il browser abbia finito di leggere l’HTML, e il tuo `document.querySelector()` non vedrà il button — perché il browser non l’ha ancora processato.

Per risolvere questo, sposta il tuo elemento `script` fuori dall’elemento `head` e posizionalo alla fine dell’elemento `body` (subito prima del tag di chiusura `</body>`).

# --hints--

Il tuo elemento `script` non dovrebbe essere dentro l’elemento `head`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
assert.notEqual(script.parentElement.tagName, "HEAD");
```

Il tuo elemento `script` dovrebbe trovarsi alla fine dell’elemento `body`.

```js
const script = document.querySelector('script[data-src$="script.js"]');
const lastChild = document.querySelector('body').lastElementChild;
assert.equal(script, lastChild);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <link rel="stylesheet" href="./styles.css">
    <title>RPG - Dragon Repeller</title>
    <script src="./script.js"></script>
  </head>
  <body>
    <div id="game">
      <div id="stats">
        <span class="stat">XP: <strong><span id="xpText">0</span></strong></span>
        <span class="stat">Health: <strong><span id="healthText">100</span></strong></span>
        <span class="stat">Gold: <strong><span id="goldText">50</span></strong></span>
      </div>
      <div id="controls">
        <button id="button1">Go to store</button>
        <button id="button2">Go to cave</button>
        <button id="button3">Fight dragon</button>
      </div>
      <div id="monsterStats"></div>
      <div id="text"></div>
    </div>

  </body>
--fcc-editable-region--
</html>
```

```js
let xp = 0;
let health = 100;
let gold = 50;
let currentWeaponIndex = 0;
let fighting;
let monsterHealth;
let inventory = ["stick"];

let button1 = document.querySelector("#button1");
```
