---
id: 66ba762af611169359d9d369
title: بناء مقال صحفي
challengeType: 25
dashedName: build-a-newspaper-article
demoType: onClick
---

# --description--

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. عيّن قيمة الجذر لـ `font-size` في مستند HTML إلى `24px`.
1. يجب أن يحتوي المستند على عنصر يحمل صنف `newspaper` ويضم جميع العناصر الأخرى بداخله.
1. يجب أن يكون للعنصر `.newspaper` قيمة `font-size` تساوي `16px` وخط `Open Sans` مع نظام بديل هو `sans-serif`.
1. داخل عنصر `.newspaper`، يجب أن تضيف سبعة عناصر على الأقل: عنصر لاسم الصحيفة يحمل صنف `name`، وعنصر لتاريخ المقال يحمل صنف `date`، وعنصر للعنوان الرئيسي يحمل صنف `headline`، وعنصر للعنوان الفرعي يحمل صنف `sub-headline`، وعنصر للمؤلف يحمل صنف `author`، وفقرتان كل واحدة تحمل صنف `text`. يجب أن تملأ جميع هذه العناصر بمعلومات مقالك.
1. يجب أن يكون للعنصر `.name` قيمة `font-size` تساوي ضعف قيمة `font-size` للجذر، ويستخدم خط `Times New Roman` مع نظام بديل هو `serif`.
1. يجب أن تجعل عناصر `.name` و `.author` جميع أحرفها كبيرة باستخدام CSS.
1. يجب أن يكون للعنصر `.headline` قيمة `font-size` تساوي ضعف قيمة `font-size` للعنصر الأب، وأن يكون عريضًا.
1. يجب أن يكون للعنصر `.sub-headline` قيمة `font-weight` تساوي `100`، وقيمة `font-size` تساوي `1.5` من قيمة `font-size` للعنصر الأب، وأن يكون مائلًا.
1. يجب أن يكون العنصر `.author` عريضًا باستخدام CSS.
1. يجب أن يكون لعناصر `.text` قيمة `text-indent` تساوي `20px`.
1. يجب أن يكون لعناصر `.text` قيمة `line-height` تساوي ضعف قيمة `font-size` للعنصر الأب.
1. يجب أن يكون الحرف الأول من كل عنصر `.text` عريضًا وبضعف حجم `font-size` للعنصر الأب. استخدم المحدِّد `::first-letter` لهذا الغرض.

# --hints--

يجب أن يكون لديك محدد `html` يضبط `font-size` إلى `24px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('html')?.fontSize, '24px');
```

يجب أن يكون لديك عنصر يحمل صنف `newspaper`.

```js
assert.exists(document.querySelector('.newspaper'));
```

يجب أن يكون لديك محدد `.newspaper` يضبط `font-size` لعناصره إلى `16px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontSize, '16px');
```

يجب أن يكون لديك محدد `.newspaper` يضبط `font-family` لعناصره إلى `'Open Sans', sans-serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper')?.fontFamily, ['"Open Sans", sans-serif', 'Open Sans, sans-serif']);
```

يجب أن يكون لديك عنصر يحمل صنف `name` داخل عنصر `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .name'));
```

يجب أن يحتوي عنصر `.name` على اسم الصحيفة الخاصة بك.

```js
assert.isAtLeast(document.querySelector('.newspaper > .name')?.innerText.length, 1);
```

يجب أن يكون لديك عنصر يحمل صنف `date` داخل عنصر `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .date'));
```

يجب أن يحتوي عنصر `.date` على التاريخ.

```js
assert.isAtLeast(document.querySelector('.newspaper > .date')?.innerText.length, 1);
```

يجب أن يكون لديك عنصر يحمل صنف `headline` داخل عنصر `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .headline'));
```

يجب أن يحتوي عنصر `.headline` على عنوان مقالك.

```js
assert.isAtLeast(document.querySelector('.newspaper > .headline')?.innerText.length, 1);
```

يجب أن يكون لديك عنصر يحمل صنف `sub-headline` داخل عنصر `.newspaper`.

```js
assert.exists(document.querySelector('.newspaper > .sub-headline'));
```

يجب أن يحتوي عنصر `.sub-headline` على العنوان الفرعي لمقالك.

```js
assert.isAtLeast(document.querySelector('.newspaper > .sub-headline')?.innerText.length, 1);
```

يجب أن يكون لديك عنصر يحمل صنف `author` داخل عنصر `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .author').length, 1);
```

يجب أن يحتوي عنصر `.author` على اسم مؤلف المقال.

```js
assert.isAtLeast(document.querySelector('.newspaper > .author')?.innerText.length, 1);
```

يجب أن يكون لديك على الأقل عنصران فقرة، كل منهما يحمل صنف `text`، داخل عنصر `.newspaper`.

```js
assert.isAtLeast(document.querySelectorAll('.newspaper > .text').length, 2);
```

يجب أن تحتوي عناصر `.text` على نص مقالك.

```js
const textEls = document.querySelectorAll('.newspaper > .text');
assert.isAtLeast(textEls.length, 1);
textEls.forEach(el => assert.isAtLeast(el.innerText.length, 1))
```

يجب أن يكون لديك محدد `.name` يضبط `font-size` لعناصره إلى `2rem`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.name')?.fontSize, '2rem');
```

يجب أن يكون لديك محدد `.name` يضبط `font-family` لعناصره إلى `'Times New Roman', serif`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.name')?.fontFamily, ['"Times New Roman", serif', 'Times New Roman, serif']);
```

يجب أن يكون لعنصر `.name` خاصية `text-transform` مضبوطة على `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.name')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

يجب أن يكون لعنصر `.author` خاصية `text-transform` مضبوطة على `uppercase`.

```js
const condition1 = new __helpers.CSSHelp(document).getStyle('.author')?.textTransform === 'uppercase';
const condition2 = new __helpers.CSSHelp(document).getStyle('.author, .name')?.textTransform === 'uppercase';
const condition3 = new __helpers.CSSHelp(document).getStyle('.name, .author')?.textTransform === 'uppercase';
assert.isTrue(condition1 || condition2 || condition3);
```

يجب أن يكون لديك محدد `.headline` يضبط `font-size` لعناصره إلى `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.headline')?.fontSize, '2em');
```

يجب أن يكون لديك محدد `.headline` يضبط `font-weight` لعناصره إلى `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.headline')?.fontWeight, ['bold', '700']);
```

يجب أن يكون لديك محدد `.sub-headline` يضبط `font-weight` لعناصره إلى `100`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontWeight, '100');
```

يجب أن يكون لديك محدد `.sub-headline` يضبط `font-size` لعناصره إلى `1.5em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontSize, '1.5em');
```

يجب أن يكون لديك محدد `.sub-headline` يضبط `font-style` لعناصره إلى `italic`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.sub-headline')?.fontStyle, 'italic');
```

يجب أن يكون لديك محدد `.author` يضبط `font-weight` لعناصره إلى `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.author')?.fontWeight, ['bold', '700']);
```

يجب أن يكون لديك محدد `.text` يضبط `text-indent` لعناصره إلى `20px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.textIndent, '20px');
```

يجب أن يكون لديك محدد `.text` يضبط `line-height` لعناصره إلى `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text')?.lineHeight, '2em');
```

يجب أن يكون لديك محدد `.text::first-letter` يضبط `font-weight` لعناصره إلى `bold`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontWeight, ['bold', '700']);
```

يجب أن يكون لديك محدد `.text::first-letter` يضبط `font-size` لعناصره إلى `2em`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.text::first-letter')?.fontSize, '2em');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Article</title>
  </head>
  <body>

  </body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Daily Chuckles</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main class="newspaper">
    <h1 class="name">Daily Chuckles</h1>
    <p class="date">July 5, 2024</p>
    <h2 class="headline">Breaking: Grandma Edna Saves Earth</h2>
    <h3 class="sub-headline">Alien Invasion Foiled by Tech-Savvy Grandma's Wi-Fi Password</h3>
    <p class="author">By Jane Doe</p>
    <p class="text">In an extraordinary twist of fate, Grandma Edna found herself at the forefront of a potential crisis
      when her clever Wi-Fi security measures unexpectedly thwarted an alien invasion. As alien spacecraft descended
      upon the town, panic spread until Grandma Edna calmly intervened, resetting her router with a cryptic passphrase
      known only to her.</p>
    <p class="text">The aliens, encountering unforeseen technological barriers, struggled to breach Grandma Edna's
      fortified network. Frustrated and bewildered, they eventually retreated to reassess their invasion strategy,
      leaving residents both relieved and amazed at Grandma Edna's resourcefulness.</p>
  </main>
</body>

</html>
```

```css
html {
  font-size: 24px;
  background-color: MintCream;
  padding-top: 100px;
}

.newspaper {
  font-family: 'Open Sans', sans-serif;
  max-width: 800px;
  margin: 20px auto;
  padding: 20px;
  background-color: white;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  font-size: 16px;
}

.name {
  font-family: 'Times New Roman', serif;
  font-size: 2rem;
  text-transform: uppercase;
  text-align: center;
  margin-bottom: 10px;
}

.date {
  font-size: 1em;
  text-align: center;
  margin-bottom: 20px;
}

.headline {
  font-size: 2em;
  font-weight: bold;
  line-height: 1.2;
  margin-bottom: 10px;
}

.sub-headline {
  font-size: 1.5em;
  font-style: italic;
  font-weight: 100;
  line-height: 1.4;
  margin-bottom: 15px;
}

.author {
  font-size: 1em;
  text-transform: uppercase;
  font-weight: bold;
  margin-bottom: 20px;
}

.text {
  text-indent: 20px;
  font-size: 1em;
  line-height: 2em;
  margin-bottom: 15px;
}

.text::first-letter {
  font-size: 2em;
  font-weight: bold;
}
```
