---
id: 6853e58d52ec64349a5760ca
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 최단 경로 알고리즘을 구현할 것입니다. 그래프의 노드들 사이에서 최단 경로를 계산하고, 그 경로도 반환하는 Python 함수를 작성하세요.

예를 들어, 도시들이 서로 다른 거리의 도로로 연결된 그래프가 주어지면, 알고리즘은 한 도시에서 다른 도시로 가는 최단 라우트를 찾습니다. 만약 City A에서 City D로 이동하려면, 알고리즘은 A ⇨ B ⇨ C ⇨ D (총 15km)가 직접 A ⇨ D (20km)로 가는 것보다 더 짧다는 것을 찾아낼 수 있습니다.

시작하려면, `INF`라는 변수를 정의하고 양의 무한대를 나타내는 `float('inf')` 값을 할당하세요. 나중에 두 노드 사이의 무한 거리를 나타낼 때 사용할 것입니다.

# --hints--

파일 상단에 `INF` 변수를 정의하세요.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

변수 `float('inf')`에 `INF`을 할당하세요.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
