---
id: 686251fb49ebc489ba60ccb5
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

그래프의 인접 행렬을 나타내기 위해 2차원 리스트를 만들어야 합니다. 이 행렬은 그래프 내 노드들 사이의 간선 가중치를 나타내는 데 사용됩니다.

`adj_matrix`라는 변수를 만들고 다음 가중치를 가진 그래프를 나타내는 2차원 리스트를 할당하세요.

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

`adj_matrix`라는 변수를 만들어야 합니다.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix`는 제공된 값을 포함하는 2차원 리스트여야 합니다.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
