---
id: 686253f1b051998ad4904e3e
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이제 세 개의 매개변수를 받는 메인 함수를 만드세요: 인접 행렬, 시작 노드, 그리고 선택적인 목표 노드입니다.

`shortest_path`라는 이름의 함수를 만들고 세 개의 매개변수 `matrix`, `start_node`, `target_node`를 받으세요. `None`에 대해 기본값으로 `target_node`를 할당하세요.

`target_node` 매개변수는 기본적으로 `None`로 설정되어 있어서, 목표 노드가 지정되지 않으면 함수가 시작 노드에서 그래프의 다른 모든 노드까지의 최단 경로를 계산해야 함을 나타냅니다.

함수 본문 안에 `pass` 문장을 일단 추가하세요.

# --hints--

`shortest_path`라는 이름의 함수를 만들어야 합니다.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

함수는 세 개의 매개변수 `matrix`, `start_node`, `target_node`를 받고, `target_node`는 기본값 `None`를 가져야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
