---
id: 68625ad34aa3f58d7f3b9b07
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

그래프의 노드 수를 저장해야 합니다. 이를 위해 인접 행렬의 길이와 일치하는 변수가 필요합니다.

`shortest_path` 함수 안에 `n` 변수를 만들고 `matrix`의 길이로 설정하세요.

# --hints--

`n`를 `matrix`의 길이로 초기화해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    pass

--fcc-editable-region--
```
