---
id: 686b78bc13d518e1267448e4
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

거리 추적 외에도 각 노드에 도달하는 실제 경로도 추적해야 합니다.

각 항목이 해당 노드에 도달하는 경로를 저장하는 리스트를 만드세요. 처음에는 각 노드의 경로가 자기 자신만 포함합니다.

리스트 내포는 리스트를 간결하게 만드는 방법입니다. 예를 들어:

```py
[x * 2 for x in range(3)]
```

`paths`의 각 `[node_no]`에 대해 `node_no`을 포함하는 리스트를 생성하는 리스트 내포를 사용해 `range(n)`라는 변수를 만드세요.

# --hints--

`paths`라는 변수가 있어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

`paths`를 리스트 내포로 초기화해야 합니다.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

리스트 내포에서 `node_no`를 순회하는 데 `range(n)`를 사용해야 합니다.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

리스트 내포는 `[node_no]`의 각 `node_no`에 대해 `range(n)`을 평가해야 합니다.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
