---
id: 686b82af323098e7a5203bf4
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

알고리즘이 실행되는 동안, 이미 방문한 노드를 추적하여 한 번 이상 처리하지 않도록 해야 합니다.

이를 위해 `visited`라는 리스트를 만들고 모든 노드에 대해 `False`로 초기화하세요.

# --hints--

`visited` 함수 안에 `shortest_path`라는 리스트가 있어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

모든 노드에 대해 `visited`를 사용하여 `False` 리스트를 `[False] * n`로 초기화해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
