---
id: 686b7848fbaa41e0c2a7efef
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`distances` 리스트를 초기화했으니, 시작 노드의 거리를 업데이트해야 합니다.

시작 노드에서 자기 자신까지의 거리는 항상 `0`이므로, `start_node` 리스트의 `distances` 인덱스에 있는 값을 `0`로 설정하세요.

# --hints--

`distances[start_node]`를 `0`로 설정해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_stmt("distances[start_node] = 0")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n

--fcc-editable-region--
```
