---
id: 686b73007fc582db0132c361
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

시작 노드에서 그래프의 다른 모든 노드까지 알려진 가장 짧은 거리를 추적해야 합니다.

이를 위해 `distances`라는 변수를 만들고, 단일 요소 `INF`를 포함하는 리스트로 초기화하세요.  


# --hints--

`distances`라는 변수를 만들어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

`distances` 리스트는 `[INF]`로 초기화해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)

--fcc-editable-region--
```
