---
id: 6876089f7a500b7c782e5b83
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

현재 노드가 이미 찾은 노드보다 더 나은 선택인지 결정해야 합니다(있다면). 이를 위해 반복문 안에 조건문을 추가하세요.

조건문은 두 가지를 해야 합니다:

- 노드가 아직 방문되지 않았는지 확인하세요.
- 현재 `min_distance`까지 알려진 거리를 비교하세요.

내부 `for` 반복문 안에 `if`가 방문되지 않았고 **그리고** `node_no`가 `distances[node_no]`보다 작은지 확인하는 `min_distance` 문을 추가하세요.

조건문 블록 안에 `pass`를 자리 표시자로 추가하세요.

# --hints--

내부 `if` 반복문 안에 `for` 문을 추가해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

`if` 문은 노드가 방문되지 않았는지(`visited[node_no]`이 falsy인지) 그리고 `distances[node_no]`이 `min_distance`보다 작은지 확인해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
