---
id: 68760c4c2648e7832c5eef16
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

가장 가까운 방문하지 않은 노드를 찾는 반복문 다음에, 실제로 유효한 노드를 찾았는지 확인해야 합니다.

만약 그런 노드가 없다면, 남은 노드들은 시작 노드에서 도달할 수 없다는 뜻이며, 알고리즘은 조기에 종료해야 합니다.

중첩된 `for` 반복문과 같은 수준에 `if` 문을 추가하여 `current == -1`인지 확인하고, 참이면 반복문을 빠져나가세요.

# --hints--

내부 `if` 반복문이 끝난 후에도 `current`가 여전히 `-1`인지 확인하는 `for` 문이 있어야 합니다.

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

반복문을 빠져나가기 위해 `break` 문을 사용하세요.

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
