---
id: 68760464eb9c2e79cc913fac
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

이번 단계에서는 그래프의 각 노드마다 한 번씩 실행되는 반복문을 추가할 것입니다. 이 반복문은 알고리즘이 여러 차례에 걸쳐 거리와 경로를 업데이트할 수 있게 합니다.

`for` 반복문을 `n`번 실행하세요. 반복 변수로는 반복 횟수 값이 필요 없으니 `_`를 사용하세요.

반복문 안에서는 다음에 처리할 노드를 선택할 준비를 위해 두 개의 변수를 만들어야 합니다:

- 현재 반복에서 발견된 가장 작은 거리를 저장할 변수  
- 그리고 이 가장 작은 거리를 가진 노드의 인덱스를 저장할 변수

`min_distance`와 `current` 변수를 만들고 각각 `INF`와 `-1`로 초기화하세요.

# --hints--

`for` 함수 안에 `shortest_path` 반복문이 있어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

`for` 반복문은 `range(n)`을 순회해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

for 반복문은 `_`를 반복 변수로 사용해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

for 반복문 안에 `min_distance`라는 변수를 만들고 `INF`로 설정해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

for 반복문 안에 `current`라는 변수를 만들고 `-1`로 설정해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
