---
id: 68760b53e250f982473e1808
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

방금 추가한 조건문이 참이라면, 현재 노드가 지금까지 찾은 최적의 방문하지 않은 선택지라는 뜻이므로 변수를 그에 맞게 업데이트해야 합니다.

`if` 블록 안에서 `min_distance`을 `distances[node_no]`로 업데이트하고 `current`을 `node_no`로 설정하세요.

# --hints--

`min_distance` 변수를 `distances[node_no]`로 업데이트해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("min_distance = distances[node_no]")`
        )
      );
    }
});
```

`current`을 `node_no`로 설정하세요.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("current = node_no")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                pass

--fcc-editable-region--
```
