---
id: 68760671d950767b055bf2f9
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

아직 방문하지 않은 노드 중에서 가장 작은 알려진 거리를 가진 노드를 찾으려면 모든 노드를 확인해야 합니다.

이를 위해 메인 루프 안에 `for` 반복문을 추가하세요.

이 반복문은 그래프의 노드 수가 `node_no`일 때 `range(n)`의 각 `n`을 순회해야 합니다. 지금은 `pass`를 반복문 본문으로 추가하세요.

# --hints--

`for` 함수 안에 두 번째 `shortest_path` 반복문이 있어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

`for` 반복문은 `range(n)`를 `node_no`를 반복 변수로 하여 순회해야 합니다.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
