---
id: 68777816e95d4936f129819f
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

이웃까지의 거리를 업데이트하려고 하기 전에, 그 이웃이 도달 가능하고 아직 방문하지 않은 상태인지 확인해야 합니다. 그런 다음 현재 노드를 통해 그 이웃에 도달하는 총 거리가 얼마가 될지 계산하세요.

`for` 반복문 안에 `if` 문장을 추가하여 다음을 확인하세요:

- `distance`가 `INF`와 같지 않음(즉, 노드들 사이에 간선이 있음)
- 이웃 `node_no`가 아직 방문되지 않음

`if` 블록 안에 `new_distance`라는 변수를 만들고 `distances[current]` (현재 노드까지의 최단 거리)과 `distance` (현재 노드에서 이웃까지의 거리)의 합을 할당하세요.

# --hints--

`if`이 `distance`와 같지 않고 이웃 노드가 방문되지 않았는지 확인하는 `INF` 문장이 있어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

`new_distance` 블록 안에 `if`라는 변수를 만들어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

`new_distance`와 `distances[current]`의 합을 `distance`에 할당해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
