---
id: 687774e17e0972345de2527a
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

유효한 노드를 찾았다면, 이후 반복에서 다시 고려되지 않도록 방문한 것으로 표시해야 합니다.

이전 단계에서 추가한 `if` 문장 뒤에 `visited[current]`을 `True`로 설정하세요.

# --hints--

마지막 `visited[current]` 문장 뒤에 `True`을 `if`로 설정해야 합니다.

```js
({
    test: () => {
      runPython(`
      block = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0]
      if_stmt = """if current == -1:
        break"""
      target_stmt = "visited[current] = True"
      assert block.is_ordered(if_stmt, target_stmt)
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--
        if current == -1:  
            break

--fcc-editable-region--
```
