---
id: 687775bed85144353d298665
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

노드를 선택하고 방문한 것으로 표시했으니, 이제 그 노드의 모든 이웃을 살펴보면서 더 짧은 경로가 있는지 확인하세요.

`visited[current] = True` 줄 다음에 `for` 루프를 추가하여 `node_no` 안의 `range(n)`를 순차적으로 반복하세요.

이 루프 안에서 변수 `distance`를 만들고 `matrix[current][node_no]`로 설정하세요. 이렇게 하면 현재 노드에서 이웃 노드까지의 거리를 알 수 있습니다.

# --hints--

현재 노드를 방문한 것으로 표시한 후 `for` 루프를 추가하세요.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 2
      `)
      );
    }
});
```

루프는 `node_no` 안의 `range(n)`를 순차적으로 반복해야 합니다.

```js
({
    test: () => {
      runPython(`  
loop = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1]  
assert loop.find_for_vars().is_equivalent("node_no")  
assert loop.find_for_iter().is_equivalent("range(n)")  
`);
    }
});
```

루프 안에서 변수 `distance`를 만들어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
neighbor_loop = outer_loop.find_bodies()[0].find_for_loops()[1]
neighbor_loop.find_bodies()[0].has_variable("distance")
      `)
      );
    }
});
```

변수 `distance`는 `matrix[current][node_no]`로 설정해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].is_equivalent("distance = matrix[current][node_no]")
      `)
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]  
                current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--

--fcc-editable-region--
```
