---
id: 687b48ccb46c5e6126d30bc6
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

마지막으로, 현재 노드에 대한 결과를 출력하세요.

출력에는 `start_node`에서 `node_no`까지의 거리와 그 경로가 모두 표시되어야 합니다.

출력 형식은 `f-string`를 사용하여 거리와 전체 경로를 다음 형식으로 보여주세요:

```md
\n[starting node]-[node number] distance: [distance]\nPath: [path]
```

여기서:

- `[starting node]`는 `start_node`입니다
- `[node number]`는 `node_no`입니다
- `[distance]`는 `start_node`에서 `node_no`까지의 거리입니다
- `[path]`는 `node_no`까지 도달한 경로입니다

# --hints--

`print(f'\n{start_node}-{node_no} distance: {distances[node_no]}\nPath: {path}')`를 사용하여 올바른 출력을 출력하세요.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function('shortest_path').find_for_loops()[1].has_call('print(f"\\\\n{start_node}-{node_no} distance: {distances[node_no]}\\\\nPath: {path}")')`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue
        string_path = (str(n) for n in paths[node_no])
        path = ' -> '.join(string_path)
        
--fcc-editable-region--
```
