---
id: 687b3c7a3af0df59b363712e
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

노드로 가는 더 짧은 경로를 찾으면, 그 노드에 도달하는 실제 경로도 업데이트해야 합니다.

같은 조건문 블록 안에서, 이웃 노드의 인덱스에 있는 `paths` 리스트를 새롭고 더 짧은 경로를 반영하도록 업데이트하세요.

`paths[node_no]`을 `current` 노드까지의 현재 경로로, 그리고 끝에 `node_no` (이웃 노드)을 추가한 경로로 업데이트해야 합니다.

# --hints--

`paths[node_no]`을 현재 경로에 이웃 노드를 더한 경로로 업데이트해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("paths[node_no] = paths[current] + [node_no]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance

--fcc-editable-region--
```
