---
id: 687b45e16f38ad5f4e4bd799
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

노드를 표시해야 한다고 결정했으니, 경로를 명확하게 출력할 수 있도록 형식을 지정해야 합니다. 이를 위해 제너레이터 표현식을 사용하세요.

제너레이터 표현식은 리스트 내포와 비슷하지만, 리스트를 만드는 대신 한 번에 하나씩 값을 생성합니다. 대괄호 `()` 대신 소괄호 `[]` 를 사용합니다. 예를 들어:

```py
numbers = [1, 2, 3]
squared = (x**2 for x in numbers)  # Generator expression
```

루프 안에서 `if` 문장 뒤에 `string_path`라는 변수를 만드세요. `paths[node_no]`에 있는 각 노드 번호를 `str()`를 사용해 문자열로 변환하는 제너레이터 표현식을 할당하세요.

제너레이터 표현식은 `n`에 있는 각 노드 번호 `paths[node_no]`를 순회해야 합니다.

# --hints--

`string_path`라는 변수를 만들어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("string_path")
      `)
      );
    }
});
```

`string_path`에 제너레이터 표현식을 할당해야 합니다.

```js
({
  test: () => {
      runPython(`
      import ast
      assert isinstance(_Node(_code).find_function("shortest_path").find_for_loops()[1].find_bodies()[0].find_variable("string_path").tree.value, ast.GeneratorExp)
      `)
 }
})
```

제너레이터 표현식은 `paths[node_no]`를 순회해야 합니다.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert path_var.find_comp_iters()[0].is_equivalent("paths[node_no]")
      `);
    }
});
```

제너레이터 표현식은 `n`를 순회하기 위해 `paths[node_no]`를 반복 변수로 사용해야 합니다.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_targets()[0].is_equivalent("n")
      `);
    }
});
```

제너레이터 표현식은 `str(n)`에 있는 각 `n`에 대해 `paths[node_no]`를 평가해야 합니다.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_expr().is_equivalent("str(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue

--fcc-editable-region--
```
