---
id: 687b3ef0c2fe185b2abc4654
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

알고리즘 실행이 끝나면, 결과를 표시할 노드를 결정해야 합니다.

특정 `target_node`가 제공되었다면, 그 노드까지의 거리와 경로만 표시하세요. 그렇지 않으면 모든 노드에 대한 결과를 표시합니다.

외부 `for _ in range(n):` 반복이 끝난 후, `targets`라는 변수를 만드세요. 조건 표현식을 사용해 `[target_node]`가 `target_node`가 아니면 `None`를 할당하고, 그렇지 않으면 `range(n)`를 할당하세요.

# --hints--

`targets`라는 변수를 만들어야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("targets")`
        )
      );
    }
});
```

`targets` 변수에는 조건 표현식이 포함되어야 합니다.

```js
({
    test: () => {
      runPython(`
        import ast
        val =_Node(_code).find_function("shortest_path").find_variable("targets").tree.value
        assert isinstance(val, ast.IfExp)
      `);
    }
});
```

조건 표현식은 `target_node is not None`를 확인해야 합니다.

```js
({
    test: () => {
        runPython(`
        compare = _Node(_Node(_code).find_function("shortest_path").find_variable("targets").tree.value.test)
        assert compare.is_equivalent("target_node is not None")
        `);
    }
});
```

대상이 제공되었을 때는 `[target_node]`를, 그렇지 않을 때는 `range(n)`를 할당해야 합니다.

```js
({
  test: () => {
      runPython(`
      import ast
      val = _Node(_code).find_function("shortest_path").find_variable("targets").tree.value
      assert _Node(val.body).is_equivalent("[target_node]")
      assert _Node(val.orelse).is_equivalent("range(n)")
    `);
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no  

        if current == -1:  
            break  
        visited[current] = True  

        for node_no in range(n):  
            distance = matrix[current][node_no]  
            if distance != INF and not visited[node_no]:  
                new_distance = distances[current] + distance  
                if new_distance < distances[node_no]:  
                    distances[node_no] = new_distance  
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--

--fcc-editable-region--
```
