---
id: 68e4c0ff9edcc92771d7549a
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

현재 `distances` 리스트에는 요소가 하나만 있습니다. 하지만 그래프의 각 노드마다 하나의 거리 값이 필요합니다.

Python에서는 리스트에 정수를 곱하면 리스트가 반복됩니다. 예를 들어, `[0] * 3`는 `[0, 0, 0]`를 만듭니다.

`distances` 리스트를 `n`에 `INF`을 곱해서 `[INF]`의 `n`개 복사본을 포함하도록 업데이트하세요.

# --hints--

`distances`를 `n`의 `INF`개 복사본을 포함하도록 업데이트해야 합니다.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
