---
id: 69046179fba35c03dec84fb6
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

이제 `join()` 메서드를 사용해 노드 번호의 문자열 표현을 하나의 읽기 쉬운 경로로 결합하세요.

`join()` 메서드는 반복 가능한 객체(제너레이터 표현식처럼)를 받아 모든 요소를 하나의 문자열로 결합하며, 각 요소 사이에 구분자를 넣습니다. 예를 들어:

```py
numbers = ['1', '2', '3']
route = ' -> '.join(numbers) # route will be '1 -> 2 -> 3'
```

`path`라는 변수를 만들고 `string_path`를 구분자로 하여 `' -> '`을 결합한 결과를 할당하세요.

# --hints--

`path`라는 변수를 만들어야 합니다.

```js
({
test: () => {
  assert(
    runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("path")
      `)
);
    }
});
```

`join`를 사용해 `string_path`의 항목들을 `' -> '` 구분자로 결합하고 결과를 `path`에 할당하세요.

```js
({
  test: () => {
    runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
assert for_loop.find_bodies()[0].find_variable("path").is_equivalent('path = " -> ".join(string_path)')
    `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no

        if current == -1:
            break
        visited[current] = True

        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]

--fcc-editable-region--
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue
        string_path = (str(n) for n in paths[node_no])
        
--fcc-editable-region--
```
