---
id: 669e2f60e83c011754f711f9
title: Crie uma página de agência de viagens
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma declaração `DOCTYPE`.
1. Você deve ter um elemento `html` com `lang` definido como `en`.
1. Você deve ter um elemento `head` contendo um elemento vazio `meta` com `charset` definido como `utf-8` e um `title` com o texto `Travel Agency Page`.
1. Você deve ter uma tag `meta` no seu elemento `head` que contenha uma descrição curta do seu site para SEO.
1. Você deve ter um elemento `h1` para apresentar os destinos de viagem.
1. Você deve ter um parágrafo abaixo do elemento `h1` introduzindo as oportunidades de viagem.
1. Você deve ter um elemento `h2` com o texto `Packages`.
1. Você deve ter um elemento `p` introduzindo brevemente os vários pacotes.
1. Você deve ter um elemento de lista não ordenada com dois itens de lista. Os dois itens de lista devem conter os textos `Viagens em Grupo` e `Passeios Privados`, respectivamente. O texto de cada item de lista deve estar envolvido por um elemento de âncora.
1. Você deve ter um elemento `h2` com o texto `Top Itineraries`.
1. Você deve ter pelo menos três elementos `figure`, cada um contendo um elemento de âncora e um elemento `figcaption`.
1. Os três elementos de âncora devem conter um elemento `img` com um atributo `alt` apropriado e um atributo `src` definido para uma imagem válida como seu conteúdo. Você pode usar `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` e `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` se desejar.
1. Todos os cinco elementos de âncora devem ter um atributo `href` com o valor `https://www.freecodecamp.org/learn` e um atributo `target` com o valor `_blank`.

# --hints--

A página da agência de viagens deve ter uma declaração `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter um elemento `html` com `lang` definido como `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Você deve ter um elemento `head` dentro do elemento `html`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Você deve ter dois elementos `meta` dentro do seu elemento `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Um elemento `meta` deve ter um atributo `name` com valor `description` e um atributo `content` não vazio.

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

Um elemento `meta` deve ter seu atributo `charset` definido como `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Você deve ter um elemento `title` dentro do elemento `head`.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

O elemento `title` deve ter o nome da agência de viagem.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Você deve ter um elemento `body` dentro do elemento `html`.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Você deve ter um elemento `h1` para apresentar os destinos de viagem.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Você deve ter apenas um elemento `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Você deve ter um elemento `p` abaixo do elemento `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

Seu primeiro elemento `p` deve apresentar as oportunidades de viagem.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

O primeiro elemento `h2` deve ter o texto `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

Você deve ter um elemento `p` abaixo do seu primeiro elemento `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Seu segundo elemento `p` deve apresentar brevemente os diversos pacotes.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Você deve ter um elemento de lista não ordenada abaixo do segundo elemento `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

Você deve ter dois itens na lista não ordenada.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Os dois itens de lista devem conter um elemento de âncora.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

O elemento de âncora do primeiro item de lista deve envolver o texto `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

O elemento de âncora do segundo item de lista deve envolver o texto `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

Você deve ter um elemento `h2` após a lista não ordenada.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

O segundo elemento `h2` deve ter o texto `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

Você deve ter pelo menos três elementos `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Cada elemento `figure` deve conter um elemento de âncora como primeiro filho.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Cada elemento `figure` deve conter um elemento `figcaption` como segundo filho.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Cada `figcaption` deve conter algum texto.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Cada um dos elementos `a` que são filhos dos elementos `figure` deve conter uma imagem.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Cada elemento `img` deve ter um atributo `src` válido.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Cada elemento `img` deve ter um atributo `alt` com um valor apropriado.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Cada elemento `a` deve ter um atributo `href` com o valor `https://www.freecodecamp.org/learn`. Não esqueça dos links nos itens da lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Cada elemento `a` deve ter um atributo `target` com o valor `_blank`. Não esqueça dos links nos itens da lista.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
