---
id: 66bba6fff611169359d9d36a
title: 비밀 이메일 페이지 만들기
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `main` 요소가 `id`가 `email`인 상태여야 합니다.
1. `#email` 요소는 `padding`가 `50px`, 위쪽 여백이 `50px`, `width`가 `500px`, 그리고 `border`가 `2px`이어야 합니다.
1. `#email` 요소의 전체 너비는 패딩과 테두리를 포함하여 `500px`여야 합니다.
1. `div` 요소 안에 id가 각각 `confidential`와 `top-secret`인 두 개의 `#email` 요소가 있어야 합니다.
1. `#confidential`와 `#top-secret` 요소는 `display`가 `inline-block`이어야 합니다.
1. `#confidential`와 `#top-secret` 요소는 `padding`, 왼쪽 여백, 그리고 `border`가 있어야 합니다.
1. `#confidential` 요소는 텍스트 `CONFIDENTIAL`를 포함해야 합니다.
1. `#top-secret` 요소는 텍스트 `TOP SECRET`를 포함해야 합니다.
1. `#confidential`와 `#top-secret` 요소는 CSS 변환을 사용해 회전되어야 합니다.
1. `#email` 요소 안에 최소 세 개의 단락 요소가 있어야 합니다.
1. 단락 요소 안에 클래스가 `span`인 최소 세 개의 `blurred` 요소가 있어야 합니다.
1. `blurred` 선택자는 CSS 필터를 사용해 `3px` 요소를 흐리게 해야 합니다.

# --hints--

`main` 요소가 `id`가 `email`인 상태여야 합니다.

```js
assert.exists(document.querySelector('main#email'));
```

`#email` 선택자가 모든 면의 `padding`를 `50px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

`#email` 선택자가 `margin-top`를 `50px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

`#email` 선택자가 `width`를 `500px`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

`#email` 요소는 `2px` 테두리를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

`#email` 요소는 `box-sizing`가 `border-box`여야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

`div` 요소 안에 `id` 요소가 `confidential`가 `#email`인 상태여야 합니다.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

`div` 요소 안에 `id` 요소가 `top-secret`가 `#email`인 상태여야 합니다.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

`#confidential` 요소는 `display`가 `inline-block`로 설정되어야 합니다.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

`#top-secret` 요소는 `display`가 `inline-block`로 설정되어야 합니다.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

`#confidential` 요소는 모든 면에 `padding`가 있어야 합니다.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

`#top-secret` 요소는 모든 면에 `padding`가 있어야 합니다.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

`#confidential` 요소는 왼쪽 여백이 있어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

`#top-secret` 요소는 왼쪽 여백이 있어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

`#confidential` 요소는 `border`가 있어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

`#top-secret` 요소는 `border`가 있어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

`#confidential` 요소의 텍스트는 `CONFIDENTIAL`이어야 합니다. 텍스트가 누락되었거나 오타가 있습니다.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

`#top-secret` 요소의 텍스트는 `TOP SECRET`이어야 합니다. 텍스트가 누락되었거나 오타가 있습니다.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

`#confidential` 요소는 CSS `transform`를 사용해 회전되어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

`#top-secret` 요소는 CSS `transform`를 사용해 회전되어야 합니다.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

`#email` 요소 안에 최소 세 개의 단락 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

단락 안에 클래스가 `span`인 최소 세 개의 `blurred` 요소가 있어야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

`.blurred` 요소는 비어 있으면 안 됩니다.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

`.blurred` 선택자가 요소의 `filter`를 `blur(3px)`로 설정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
