---
id: 690ddb5e6ee94eb73ed172b0
title: 寄付フォームのデバッグ
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

地域のチャリティー団体が寄付フォームのウェブサイトを作成しましたが、いくつか修正が必要な問題があります。フォームはアクセシブルではなく、HTMLの構文エラーもあります。

すべてのエラーを修正して、フォームが正しく動作し、すべてのユーザーがアクセスできるようにしてください。以下のユーザーストーリーの項目を完了し、「Run the Tests」をクリックしてすべてのエラーが修正されたか確認してください。

**ユーザーストーリー:**

1. `input` 要素は空要素であり、終了タグを持つべきではありません。フォームからすべての `</input>` 終了タグを削除してください。
2. 各フォームの入力フィールドに対して `label` 要素を追加し、ユーザーが各フィールドの目的を理解できるようにしてください。ラベルのテキストは現在各入力の隣にあるものと一致させてください。
3. `Email Address:` の入力タイプは `email` ではなく `text` にする必要があります。
4. 各 `label` 要素を対応する `input` 要素と `for` 属性（`label` に付ける）および一致する `id` 属性（`input` に付ける）を使って関連付けてください。
5. テキスト、メール、数値の入力フィールド（チェックボックスや送信ボタンは除く）に `required` 属性を追加し、ユーザーが必須情報を入力するようにしてください。

# --hints--

コード内に `</input>` 終了タグがあってはいけません。

```js
assert.notMatch(code, /<\/input>/);
```

フォームには正確に5つの `input` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

フォームには正確に4つの `label` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

最初の `label` はテキスト `Full Name:` を持つべきです。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

最初の `input` は `required` 属性を持つべきです。

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

2番目の `label` はテキスト `Email Address:` を持つべきです。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

2番目の `input` は `required` 属性を持つべきです。

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

3番目の `label` はテキスト `Donation Amount ($):` を持つべきです。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

3番目の `input` は `required` 属性を持つべきです。

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

4番目の `label` はテキスト `Subscribe` を持つべきです。

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

すべての `input` 要素は `label` 要素と `for` および `id` 属性を使って関連付けるべきです。

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

メールアドレス用の `input` はタイプ `email` であるべきです。

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

チェックボックスの `input` は `required` 属性を持ってはいけません。

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

送信ボタンの `button` は `required` 属性を持ってはいけません。

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
