---
id: 61764f7e5240eb7ccc7f6a0a
title: ステップ 13
challengeType: 0
dashedName: step-13
---

# --description--

次に、クラス `one` を選択する新しい CSS ルールを作成し、その `background-color` プロパティを `red` に設定してください。

# --hints--

クラス `one` を選択するクラスセレクターを使用する必要があります。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.one'));
```

Your `.one` CSS rule should have a `background-color` property set to `red`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor, 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--

--fcc-editable-region--

```
