---
id: 619b7424f43ec9215e538afe
title: ステップ 39
challengeType: 0
dashedName: step-39
---

# --description--

Next, change the class `two` to `green` in the second marker `div` and change the class `three` to `blue` in the third marker `div`.

# --hints--

2 つ目のマーカーの `div` にクラス `two` は必要ありません。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isFalse(containerSecondChild?.classList?.contains('two'));
```

2 つ目のマーカーの `div` には `marker` と `green` というクラスが必要です。

```js
const containerSecondChild = [...document.querySelector('.container')?.children][1];
assert.isTrue(containerSecondChild?.classList?.contains('marker'));
assert.isTrue(containerSecondChild?.classList?.contains('green'));
```

3 つ目のマーカーの `div` にクラス `three` は必要ありません。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isFalse(containerThirdChild?.classList?.contains('three'));
```

3 つ目のマーカーの `div` には `marker` と `blue` というクラスが必要です。

```js
const containerThirdChild = [...document.querySelector('.container')?.children][2];
assert.isTrue(containerThirdChild?.classList?.contains('marker'));
assert.isTrue(containerThirdChild?.classList?.contains('blue'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      --fcc-editable-region--
      <div class="marker red">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
      --fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.red {
  background-color: rgb(255, 0, 0);
}

.two {
  background-color: rgb(0, 0, 0);
}

.three {
  background-color: rgb(0, 0, 0);
}

```
