---
id: 5a24c314108439a4d4036150
title: Обробка дії в store
challengeType: 6
forumTopicId: 301444
dashedName: handle-an-action-in-the-store
---

# --description--

Після створення та відправлення дії, Redux store має знати, як на неї реагувати. Це завдання `reducer` функції. Редюсери в Redux відповідають за зміни стану, які відбуваються у відповідь на дії. `reducer` приймає `state` та `action` як аргументи і завжди повертає новий `state`. Важливо розуміти, що це **єдина** роль редюсера. Він не має побічних ефектів — ніколи не викликає API кінцеву точку і не має прихованих сюрпризів. Редюсер — це просто чиста функція, яка приймає стан і дію, а потім повертає новий стан.

Ще один ключовий принцип у Redux полягає в тому, що `state` є лише для читання. Іншими словами, `reducer` функція **завжди** має повертати нову копію `state` і ніколи не змінювати стан безпосередньо. Redux не примушує до незмінності стану, проте ви несете відповідальність за її дотримання у коді своїх редюсерів. Ви попрактикуєте це у наступних завданнях.

# --instructions--

У редакторі коду є попередній приклад, а також початок `reducer` функції для вас. Заповніть тіло `reducer` функції так, щоб якщо вона отримує дію типу `'LOGIN'`, вона повертала об’єкт стану з `login`, встановленою в `true`. В іншому випадку вона має повертати поточний `state`. Зверніть увагу, що поточний `state` і відправлена `action` передаються редюсеру, тож ви можете безпосередньо отримати тип дії через `action.type`.

# --hints--

Виклик функції `loginAction` має повертати об’єкт з властивістю type, встановленою в рядок `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

Store має бути ініціалізований об’єктом з властивістю `login`, встановленою в `false`.

```js
assert(store.getState().login === false);
```

Відправлення `loginAction` має оновити властивість `login` у стані store до `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginAction());
    const afterState = store.getState();
    return initialState.login === false && afterState.login === true;
  })()
);
```

Якщо дія не має типу `LOGIN`, store має повертати поточний стан.

```js
assert(
  (function () {
    store.dispatch({ type: '__TEST__ACTION__' });
    let afterTest = store.getState();
    return typeof afterTest === 'object' && afterTest.hasOwnProperty('login');
  })()
);
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```

# --solutions--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {

  if (action.type === 'LOGIN') {
    return {login: true}
  }

  else {
    return state
  }

};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```
