---
id: 5a24c314108439a4d4036158
title: Ніколи не змінюйте стан
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

Ці фінальні завдання описують кілька способів забезпечення ключового принципу незмінності стану в Redux. Незмінний стан означає, що ви ніколи не змінюєте стан безпосередньо, натомість повертаєте нову копію стану.

Якщо зробити знімок стану застосунку Redux у часі, ви побачите щось на кшталт `state 1`, `state 2`, `state 3`, `state 4`, `...` і так далі, де кожен стан може бути схожим на попередній, але кожен є окремим фрагментом даних. Саме ця незмінність забезпечує такі функції, як налагодження з можливістю подорожі в часі, про які ви, можливо, чули.

Redux не примушує активно забезпечувати незмінність стану у своєму сховищі або редюсерах, ця відповідальність лежить на програмісті. На щастя, JavaScript (особливо ES6) надає кілька корисних інструментів, які ви можете використати, щоб забезпечити незмінність свого стану, незалежно від того, чи це `string`, `number`, `array` або `object`. Зверніть увагу, що рядки та числа — це примітивні значення і за своєю природою незмінні. Іншими словами, 3 завжди 3. Ви не можете змінити значення числа 3. Однак `array` або `object` є змінними. На практиці ваш стан, ймовірно, складатиметься з `array` або `object`, оскільки це корисні структури даних для представлення багатьох типів інформації.

# --instructions--

У редакторі коду є `store` і `reducer` для керування елементами списку справ. Завершіть написання випадку `ADD_TO_DO` у редюсері, щоб додати нову справу до стану. Існує кілька способів зробити це за допомогою стандартного JavaScript або ES6. Спробуйте знайти спосіб повернути новий масив із доданим в кінець елементом із `action.todo`.

# --hints--

Сховище Redux має існувати і ініціалізуватися зі станом, рівним масиву `todos` у редакторі коду.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo` і `immutableReducer` обидва мають бути функціями.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Відправлення дії типу `ADD_TO_DO` у сховище Redux має додати елемент `todo` і НЕ має змінювати стан.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
