---
id: 5a24c314108439a4d4036157
title: Напишіть фішку з Redux
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

Тепер ви вивчили всі основні принципи Redux! Ви побачили, як створювати дії та генератори дій, створювати Redux store, відправляти свої дії до store і проєктувати оновлення стану за допомогою чистих редюсерів. Ви навіть бачили, як керувати складним станом за допомогою композиції редюсерів і обробляти асинхронні дії. Ці приклади прості, але ці концепції є основними принципами Redux. Якщо ви добре їх розумієте, ви готові почати створювати свій власний застосунок на Redux. Наступні завдання охоплюють деякі деталі щодо `state` незмінності, але спершу ось повторення всього, що ви вивчили досі.

# --instructions--

У цьому уроці ви реалізуєте просту фішку з Redux з нуля. Основи надані в редакторі коду, але вам доведеться заповнити деталі! Використовуйте надані імена та визначте генератори дій `incAction` і `decAction`, типи дій `counterReducer()`, `INCREMENT` і `DECREMENT`, а також фінальний Redux `store`. Коли ви закінчите, ви зможете відправляти дії `INCREMENT` або `DECREMENT`, щоб збільшувати або зменшувати стан, що зберігається в `store`. Успіхів у створенні вашого першого застосунку на Redux!

# --hints--

Генератор дій `incAction` має повертати об’єкт дії з `type`, рівним значенню `INCREMENT`

```js
assert(incAction().type === INCREMENT);
```

Генератор дій `decAction` має повертати об’єкт дії з `type`, рівним значенню `DECREMENT`

```js
assert(decAction().type === DECREMENT);
```

Виклик `store.getState()` має повертати число

```js
assert(typeof store.getState() === 'number');
```

Redux store має ініціалізуватися зі станом `state`, рівним 0.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Відправлення `incAction` у Redux store має збільшувати `state` на 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Відправлення `decAction` у Redux store має зменшувати `state` на 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` має бути функцією

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
