---
id: 5a24c314108439a4d4036152
title: Використовуйте const для типів дій
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Загальною практикою при роботі з Redux є призначення типів дій як констант лише для читання, а потім посилання на ці константи там, де вони використовуються. Ви можете рефакторити код, з яким працюєте, щоб записати типи дій як `const` оголошення.

# --instructions--

Оголосіть `LOGIN` і `LOGOUT` як `const` значення та призначте їм рядки `'LOGIN'` і `'LOGOUT'` відповідно. Потім відредагуйте `authReducer()` і творці дій, щоб вони посилалися на ці константи замість рядкових значень.

**Примітка:** Загалом існує конвенція писати константи великими літерами, і це також стандартна практика в Redux.

# --hints--

Виклик функції `loginUser` має повертати об’єкт із властивістю `type`, призначеною рядку `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Виклик функції `logoutUser` має повертати об’єкт із властивістю `type`, призначеною рядку `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

Сховище має бути ініціалізоване об’єктом із властивістю `login`, призначеною `false`.

```js
assert(store.getState().authenticated === false);
```

Відправлення `loginUser` має оновити властивість `login` у стані сховища на `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Відправлення `logoutUser` має оновити властивість `login` у стані сховища на `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Функція `authReducer` має обробляти кілька типів дій за допомогою інструкції switch.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN` і `LOGOUT` мають бути оголошені як `const` значення і призначені рядки `LOGIN` і `LOGOUT`.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

Творці дій і редуктор мають посилатися на константи `LOGIN` і `LOGOUT`.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
