---
id: 5a24c314108439a4d4036159
title: Використання оператора розширення для масивів
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Одним із рішень у ES6 для забезпечення незмінності стану в Redux є оператор розширення: `...`. Оператор розширення має різноманітні застосування, одне з яких добре підходить для попереднього завдання — створення нового масиву на основі існуючого. Це відносно новий, але широко використовуваний синтаксис. Наприклад, якщо у вас є масив `myArray` і ви напишете:

```js
let newArray = [...myArray];
```

`newArray` тепер є копією `myArray`. Обидва масиви існують окремо в пам’яті. Якщо ви виконаєте мутацію, як `newArray.push(5)`, `myArray` не зміниться. `...` фактично *розгортає* значення з `myArray` у новий масив. Щоб клонувати масив, але додати додаткові значення у новий масив, ви можете написати `[...myArray, 'new value']`. Це поверне новий масив, що складається зі значень `myArray` і рядка `new value` як останнього елемента. Синтаксис розширення можна використовувати кілька разів при формуванні масиву, але важливо пам’ятати, що він робить лише поверхневу копію масиву. Тобто він забезпечує незмінні операції масиву лише для одномірних масивів.

# --instructions--

Використайте оператор розширення, щоб повернути нову копію стану, коли додається завдання.

# --hints--

Redux store має існувати і ініціалізуватися зі станом, рівним `["Do not mutate state!"]`.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo` і `immutableReducer` мають бути функціями.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Відправлення дії типу `ADD_TO_DO` у Redux store має додати елемент `todo` і НЕ мутувати стан.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

Оператор розширення має бути використаний для повернення нового стану.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
