---
id: 5a24c314108439a4d4036153
title: Зареєструвати слухача сховища
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Ще один метод, до якого ви маєте доступ у об’єкті Redux `store`, це `store.subscribe()`. Він дозволяє підписувати функції-слухачі на сховище, які викликаються щоразу, коли до сховища надсилається дія. Простий приклад використання цього методу — підписати функцію на ваше сховище, яка просто виводить повідомлення щоразу, коли отримується дія і стан сховища оновлюється.

# --instructions--

Напишіть функцію зворотного виклику, яка збільшує глобальну змінну `count` щоразу, коли сховище отримує дію, і передайте цю функцію у метод `store.subscribe()`. Ви побачите, що `store.dispatch()` викликається тричі поспіль, кожного разу безпосередньо передаючи об’єкт дії. Слідкуйте за вихідними даними в консолі між відправленнями дій, щоб побачити, як відбуваються оновлення.

# --hints--

Відправлення дії `ADD` у сховище має збільшувати стан на `1`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

До сховища має бути підписана функція-слухач за допомогою `store.subscribe`.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` має отримати функцію.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

Функція, передана у `store.subscribe`, не повинна викликатися.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

Функція зворотного виклику для `store.subscribe` також має збільшувати глобальну змінну `count` під час оновлення сховища.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
