---
id: 5a24c314108439a4d403614f
title: Відправлення події дії
challengeType: 6
forumTopicId: 301442
dashedName: dispatch-an-action-event
---

# --description--

Метод `dispatch` використовується для відправлення дій у Redux store. Виклик `store.dispatch()` з передачею значення, що повертається з action creator, відправляє дію назад у store.

Нагадаємо, що action creators повертають об’єкт із властивістю type, яка вказує тип дії, що відбулася. Потім метод відправляє об’єкт дії у Redux store. Виходячи з прикладу попереднього завдання, наступні рядки є еквівалентними, і обидва відправляють дію типу `LOGIN`:

```js
store.dispatch(actionCreator());
store.dispatch({ type: 'LOGIN' });
```

# --instructions--

Redux store у редакторі коду має ініціальний стан, який є об’єктом із властивістю `login`, що наразі встановлена в `false`. Також є action creator під назвою `loginAction()`, який повертає дію типу `LOGIN`. Відправте дію `LOGIN` у Redux store, викликавши метод `dispatch` і передавши дію, створену `loginAction()`.

# --hints--

Виклик функції `loginAction` має повернути об’єкт із властивістю `type`, встановленою в рядок `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

Store має бути ініціалізований об’єктом із властивістю `login`, встановленою в `false`.

```js
assert(store.getState().login === false);
```

Для відправлення дії типу `store.dispatch()` слід використовувати метод `LOGIN`.

```js
assert(
    (function () {
      let noWhiteSpace = code.replace(/\s/g, '');
      return (
        noWhiteSpace.includes('store.dispatch(loginAction())') ||
        noWhiteSpace.includes("store.dispatch({type: 'LOGIN'})") === true
      );
    })()
);
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

// Dispatch the action here:
```

# --solutions--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

store.dispatch(loginAction());
```
