---
id: 5a24c314108439a4d403614e
title: визначте creator дії
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

Після створення дії наступним кроком є відправлення цієї дії до Redux store, щоб він міг оновити свій стан. У Redux для цього визначають creator дій. Creator дії — це просто JavaScript функція, яка повертає дію. Іншими словами, creator дій створюють об’єкти, які представляють події дії.

# --instructions--

Визначте функцію з іменем `actionCreator()`, яка повертає об’єкт `action` при виклику.

# --hints--

Функція `actionCreator` має існувати.

```js
assert(typeof actionCreator === 'function');
```

Виклик функції `actionCreator` має повертати об’єкт `action`.

```js
assert(typeof action === 'object');
```

Повернутий `action` має містити ключову властивість `type` зі значенням `LOGIN`.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
