---
id: 5a24c314108439a4d403614d
title: оголосити дію Redux
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Оскільки Redux — це фреймворк для керування станом, оновлення стану є однією з його основних задач. У Redux усі оновлення стану відбуваються через відправлення дій. Дія — це просто JavaScript-об’єкт, який містить інформацію про подію дії, що сталася. Сховище Redux отримує ці об’єкти дій і відповідно оновлює свій стан. Іноді дія Redux також містить деякі дані. Наприклад, дія може містити ім’я користувача після входу користувача в систему. Хоча дані є необов’язковими, дії мають містити властивість `type`, яка вказує тип дії, що відбулася.

Уявіть дії Redux як посланців, які доставляють інформацію про події, що відбуваються у вашому застосунку, до сховища Redux. Потім сховище виконує оновлення стану на основі отриманої дії.

# --instructions--

Написати дію Redux так само просто, як оголосити об’єкт із властивістю type. Оголосіть об’єкт `action` і надайте йому властивість `type` зі значенням рядка `'LOGIN'`.

# --hints--

Повинен існувати об’єкт `action`.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

Об’єкт `action` має мати ключову властивість `type` зі значенням `LOGIN`.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
