---
id: 5a24c314108439a4d4036151
title: Використання інструкції switch для обробки кількох дій
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

Ви можете вказати Redux store, як обробляти кілька типів дій. Припустімо, ви керуєте автентифікацією користувача у своєму Redux store. Ви хочете мати представлення стану для випадків, коли користувачі увійшли в систему, і коли вони вийшли. Ви представляєте це одним об’єктом стану з властивістю `authenticated`. Також вам потрібні творці дій, які створюють дії, що відповідають входу користувача та виходу користувача, разом із самими об’єктами дій.

# --instructions--

У редакторі коду для вас налаштовані store, дії та творці дій. Заповніть функцію `reducer`, щоб обробляти кілька дій автентифікації. Використайте JavaScript інструкцію `switch` у `reducer`, щоб реагувати на різні події дій. Це стандартний шаблон при написанні редюсерів Redux. Інструкція switch має переключатися за `action.type` і повертати відповідний стан автентифікації.

**Примітка:** На цьому етапі не хвилюйтеся про незмінність стану, оскільки в цьому прикладі вона невелика і проста. Для кожної дії ви можете повертати новий об’єкт — наприклад, `{authenticated: true}`. Також не забудьте написати випадок `default` у вашій інструкції switch, який повертає поточний `state`. Це важливо, тому що коли у вашому застосунку є кілька редюсерів, вони всі виконуються щоразу, коли відбувається відправлення дії, навіть якщо дія не стосується цього редюсера. У такому випадку ви хочете переконатися, що повертаєте поточний `state`.

# --hints--

Виклик функції `loginUser` має повертати об’єкт із властивістю type, встановленою в рядок `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Виклик функції `logoutUser` має повертати об’єкт із властивістю type, встановленою в рядок `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

Store має бути ініціалізований об’єктом із властивістю `authenticated`, встановленою в `false`.

```js
assert(store.getState().authenticated === false);
```

Відправлення `loginUser` має оновити властивість `authenticated` у стані store на `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Відправлення `logoutUser` має оновити властивість `authenticated` у стані store на `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

Функція `authReducer` має обробляти кілька типів дій за допомогою інструкції `switch`.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
