---
id: 5a24c314108439a4d4036156
title: Використання підпрограмного забезпечення для обробки асинхронних дій
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

До цього часу ці завдання уникали обговорення асинхронних дій, але вони є невід’ємною частиною веб-розробки. Рано чи пізно вам доведеться робити виклики асинхронних кінцевих точок у вашому Redux-застосунку, тож як обробляти такі запити? Redux надає підпрограмне забезпечення, спеціально розроблене для цієї мети, яке називається Redux Thunk middleware. Ось короткий опис, як його використовувати з Redux.

Щоб включити Redux Thunk middleware, ви передаєте його як аргумент до `Redux.applyMiddleware()`. Ця інструкція потім передається як другий необов’язковий параметр функції `createStore()`. Подивіться на код у нижній частині редактора, щоб це побачити. Потім, щоб створити асинхронну дію, ви повертаєте функцію у творці дії, яка приймає `dispatch` як аргумент. Усередині цієї функції ви можете викликати dispatch для дій і виконувати асинхронні запити.

У цьому прикладі асинхронний запит імітується викликом `setTimeout()`. Зазвичай перед початком асинхронної поведінки викликають дію, щоб стан застосунку знав, що деякі дані запитуються (цей стан, наприклад, може відображати іконку завантаження). Потім, коли ви отримуєте дані, викликаєте іншу дію, яка несе ці дані як корисне навантаження разом з інформацією про завершення дії.

Пам’ятайте, що ви передаєте `dispatch` як параметр до цього спеціального творця дії. Саме його ви використовуєте для виклику dispatch ваших дій, просто передаєте дію безпосередньо в dispatch, а підпрограмне забезпечення робить усе інше.

# --instructions--

Напишіть обидва виклики dispatch у творці дії `handleAsync()`. Викличте dispatch `requestingData()` перед `setTimeout()` (імітація виклику API). Потім, після отримання (умовних) даних, викличте dispatch дії `receivedData()`, передаючи ці дані. Тепер ви знаєте, як обробляти асинхронні дії в Redux. Все інше працюватиме, як і раніше.

# --hints--

Творець дії `requestingData` має повертати об’єкт типу, що дорівнює значенню `REQUESTING_DATA`.

```js
assert(requestingData().type === REQUESTING_DATA);
```

Творець дії `receivedData` має повертати об’єкт типу, що дорівнює значенню `RECEIVED_DATA`.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer` має бути функцією.

```js
assert(typeof asyncDataReducer === 'function');
```

Виклик dispatch дії `requestingData` має оновити властивість fetching у сховищі `state` на `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

Виклик dispatch `handleAsync` має спочатку викликати дію запиту даних, а потім після затримки викликати дію отримання даних.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
