---
id: 5a24c314108439a4d4036154
title: Об’єднання кількох редюсерів
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

Коли стан вашого застосунку починає ставати складнішим, може виникнути бажання розділити стан на кілька частин. Проте пам’ятайте перший принцип Redux: увесь стан застосунку зберігається в одному об’єкті стану у сховищі. Тому Redux пропонує композицію редюсерів як рішення для складної моделі стану. Ви визначаєте кілька редюсерів для обробки різних частин стану вашого застосунку, а потім об’єднуєте ці редюсери в один кореневий редюсер. Кореневий редюсер передається у метод Redux `createStore()`.

Щоб дозволити нам об’єднувати кілька редюсерів, Redux надає метод `combineReducers()`. Цей метод приймає об’єкт як аргумент, у якому ви визначаєте властивості, що асоціюють ключі з конкретними функціями редюсерів. Ім’я, яке ви даєте ключам, буде використане Redux як ім’я для відповідної частини стану.

Зазвичай хорошою практикою є створення редюсера для кожної частини стану застосунку, якщо вони відрізняються або є унікальними якимось чином. Наприклад, у застосунку для нотаток з автентифікацією користувача один редюсер може обробляти автентифікацію, а інший — текст і нотатки, які користувач відправляє. Для такого застосунку ми могли б написати метод `combineReducers()` так:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

Тепер ключ `notes` міститиме увесь стан, пов’язаний із нашими нотатками, і оброблюваний нашим `notesReducer`. Ось як можна об’єднати кілька редюсерів для керування складнішим станом застосунку. У цьому прикладі стан, що зберігається у Redux-сховищі, буде одним об’єктом, який містить властивості `auth` і `notes`.

# --instructions--

У редакторі коду надано функції `counterReducer()` і `authReducer()`, а також Redux-сховище. Завершіть написання функції `rootReducer()`, використовуючи метод `Redux.combineReducers()`. Призначте `counterReducer` ключу з ім’ям `count`, а `authReducer` — ключу з ім’ям `auth`.

# --hints--

`counterReducer` має збільшувати і зменшувати `state`.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

`authReducer` має перемикати стан `state` об’єкта `authenticated` між `true` і `false`.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

Сховище `state` має мати два ключі: `count`, який містить число, і `auth`, який містить об’єкт. Об’єкт `auth` має мати властивість `authenticated`, яка містить булеве значення.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

`rootReducer` має бути функцією, що об’єднує `counterReducer` і `authReducer`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
