---
id: 5a24c314108439a4d403614c
title: Отримати стан із Redux store
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Об’єкт Redux store надає кілька методів, які дозволяють взаємодіяти з ним. Наприклад, ви можете отримати поточний `state` у об’єкті Redux store за допомогою методу `getState()`.

# --instructions--

Код із попереднього завдання переписано більш лаконічно в редакторі коду. Використайте `store.getState()`, щоб отримати `state` із `store`, і призначте це новій змінній `currentState`.

# --hints--

Redux store має мати значення 5 для ініціального стану.

```js
assert(store.getState() === 5);
```

Змінна `currentState` має існувати і бути призначеною поточному стану Redux store.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
