---
id: 5a24c314108439a4d4036155
title: Надіслати дані дії до сховища
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

Ви вже навчилися відправляти дії до Redux-сховища, але поки що ці дії не містили жодної інформації, окрім `type`. Ви також можете надсилати конкретні дані разом із своїми діями. Насправді це дуже поширено, бо дії зазвичай виникають через взаємодію користувача і часто несуть із собою певні дані. Redux-сховище часто має знати про ці дані.

# --instructions--

У редакторі коду визначено базовий `notesReducer()` та `addNoteText()` творець дій. Завершіть тіло функції `addNoteText()`, щоб вона повертала об’єкт `action`. Об’єкт має містити властивість `type` зі значенням `ADD_NOTE`, а також властивість `text`, встановлену в дані `note`, які передаються у творець дії. Коли ви викликаєте творця дії, ви передаєте конкретну інформацію про нотатку, до якої можна отримати доступ для об’єкта.

Далі завершіть написання інструкції `switch` у `notesReducer()`. Вам потрібно додати випадок, який обробляє дії `addNoteText()`. Цей випадок має спрацьовувати щоразу, коли є дія типу `ADD_NOTE`, і повертати властивість `text` вхідного `action` як новий `state`.

Дія відправляється внизу коду. Коли закінчите, запустіть код і спостерігайте за консоллю. Ось і все, що потрібно, щоб надіслати специфічні для дії дані до сховища і використати їх при оновленні стану сховища `state`.

# --hints--

Творець дії `addNoteText` має повертати об’єкт із ключами `type` та `text`.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

Відправлення дії типу `ADD_NOTE` за допомогою творця дії `addNoteText` має оновити `state` на рядок, переданий у творця дії.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
