---
id: 5a24c314108439a4d403615b
title: Копіювання об’єкта за допомогою Object.assign
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

Останні кілька завдань працювали з масивами, але існують способи допомогти забезпечити незмінність стану, коли стан — це також `object`. Корисним інструментом для роботи з об’єктами є утиліта `Object.assign()`. `Object.assign()` приймає цільовий об’єкт і вихідні об’єкти та спроєктовує властивості з вихідних об’єктів на цільовий об’єкт. Будь-які співпадаючі властивості перезаписуються властивостями з вихідних об’єктів. Ця поведінка часто використовується для створення поверхневих копій об’єктів шляхом передачі порожнього об’єкта як першого аргументу, а потім об’єкта(ів), які ви хочете скопіювати. Ось приклад:

```js
const newObject = Object.assign({}, obj1, obj2);
```

Це створює `newObject` як новий `object`, який містить властивості, що наразі існують у `obj1` та `obj2`.

# --instructions--

Стан Redux і дії були змінені, щоб обробляти `object` для `state`. Відредагуйте код, щоб повертати новий об’єкт `state` для дій з типом `ONLINE`, які встановлюють властивість `status` у рядок `online`. Спробуйте використати `Object.assign()`, щоб виконати завдання.

# --hints--

Сховище Redux має існувати і ініціалізуватися зі станом, який є еквівалентним об’єкту `defaultState`, оголошеному на лінії 1.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` і `immutableReducer` обидва мають бути функціями.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

Відправлення дії типу `ONLINE` має оновити властивість `status` у стані на `online` і НЕ має мутувати стан.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

`Object.assign` має використовуватися для повернення нового стану.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
