---
id: 60f02e7361b68405e27b62a5
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

All'interno del `body`, fornisci un contesto di intestazione per il contenuto, aggiungendo un `h1` con il testo `Registration Form`.

# --hints--

Dovresti aggiungere il `h1` all'interno del `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Dovresti dare al `h1` il testo `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
