---
id: 60f027c87bc98f050395c139
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Successivamente, aggiungi i tag `head` di apertura e `body` di chiusura all'interno dell'elemento `html`.

# --hints--

Dovresti avere un tag `head` di apertura.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag `head` di chiusura.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag `body` di apertura.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag `body` di chiusura.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">

</html>
--fcc-editable-region--
```
