---
id: 60f0286404aefb0562a4fdf9
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Aggiungi un elemento `title` e un elemento `meta` all'interno dell'elemento `head`. Dai al tuo progetto un titolo `Registration Form` e aggiungi l'attributo `charset` con valore `utf-8` al tuo elemento `meta`.

# --hints--

Il tuo codice dovrebbe avere un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Il tuo elemento `title` dovrebbe essere all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Il tuo progetto dovrebbe avere un titolo `Registration Form`. Ricorda che maiuscole e minuscole e l'ortografia sono importanti per il titolo.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Dovresti creare un elemento `meta` all'interno dell'elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Dovresti dare all'elemento `meta` un attributo `charset` con valore `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Il tuo elemento `meta` dovrebbe essere un elemento vuoto, non ha un tag di chiusura `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
