---
id: 60f027099a15b00485563dd2
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Sotto il `DOCTYPE`, aggiungi un elemento `html` con un attributo `lang` impostato su `en`, così avrai un posto dove iniziare a inserire del codice.

# --hints--

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Il tuo elemento `html` dovrebbe avere un tag di apertura e un attributo `lang` con valore `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

I tuoi tag `html` dovrebbero essere nel corretto ordine.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Dovresti avere solo un elemento `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
