---
id: 60f1922fcbd2410527b3bd89
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--
L'unità `vh` indica l'altezza della viewport ed è uguale all'1% della `height` della viewport. Questo la rende relativa all'altezza della viewport.

È ora di abbellire il progetto con un po' di CSS. Inizia dando al `body` un `width` di `100%` e un `height` di `100vh`.

# --hints--

Dovresti usare il selettore elemento `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Dovresti dare al `body` un `width` di `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Dovresti dare al `body` un `height` di `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
