---
id: 60f1a9cbd23023082e149fee
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Così va meglio. Ora, rendi lo sfondo più gradevole per gli occhi, cambiando il `body` `background-color` in `#1b1b32`. Poi, per vedere il testo, cambia il `color` in `#f5f6f7`.

# --hints--

Dovresti aggiungere il `background-color` all'interno del selettore dell'elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor);
```

Dovresti impostare il `background-color` con un valore di `#1b1b32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor, 'rgb(27, 27, 50)');
```

Dovresti aggiungere il `color` all'interno del selettore dell'elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.color);
```

Dovresti impostare il `color` con un valore di `#f5f6f7`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.color, 'rgb(245, 246, 247)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
  margin: 0;
}
--fcc-editable-region--
```
