---
id: 60f1a5e2d2c23707a4f9a660
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora, elimina la barra di scorrimento orizzontale impostando il `body` predefinito `margin` aggiunto da alcuni browser su `0`.

# --hints--

Dovresti aggiungere `margin` all'interno del selettore dell'elemento `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('body')?.margin);
```

Dovresti assegnare al `margin` un valore di `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  width: 100%;
  height: 100vh;
}
--fcc-editable-region--
```
