---
id: 60f5c3e399ff1a05629964e4
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Come suggerisce il titolo, stai creando un modulo. Quindi, dopo l'elemento `p`, inserisci un elemento `form` con un attributo `action` che ha come obiettivo `https://register-demo.freecodecamp.org`.

# --hints--

Dovresti aggiungere un elemento `form` accanto all'elemento `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Dovresti dare all'elemento `form` un attributo `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Dovresti assegnare all'elemento `action` un valore `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

Il tuo elemento `form` dovrebbe avere un tag di chiusura `</form>`.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
