---
id: 60fad1cafcde010995e15306
title: Passaggio 42
challengeType: 0
dashedName: step-42
---

# --description--

Con gli invii di modulo, è utile e buona pratica fornire a ogni elemento inviabile un attributo `name`. Questo attributo serve a identificare l'elemento nell'invio del modulo.

A eccezione dei due input `radio` (che hai già nominato), assegna a ogni elemento inviabile un attributo `name` unico a tua scelta.

# --hints--

Dovresti dare all'input `input` che si aspetta un nome un attributo `name`. _PS Avrei scelto `first-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

Dovresti dare all'input `input` che si aspetta un cognome un attributo `name`. _PS Avrei scelto `last-name`_

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

Dovresti dare all'input `email` un attributo `name`. _PS Avrei scelto `email`_

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

Dovresti dare all'input `password` un attributo `name`. _PS Avrei scelto `password`_

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

Dovresti dare all'input `checkbox` un attributo `name`. _PS Avrei scelto `terms`_

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

Dovresti dare all'input `file` un attributo `name`. _PS Avrei scelto `file`_

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

Dovresti dare all'input `number` un attributo `name`. _PS Avrei scelto `age`_

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

Dovresti dare all'elemento `select` un attributo `name`. _PS Avrei scelto `referrer`_

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

Dovresti dare all'elemento `textarea` un attributo `name`. _PS Avrei scelto `bio`_

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

Non dovresti dare a nessuno degli elementi `option` un attributo `name`.

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

Non dovresti dare a nessuno degli elementi `label` un attributo `name`.

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

Non dovresti dare a nessuno degli elementi `fieldset` un attributo `name`.

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
        <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset>
        <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
        <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
        <label for="referrer">How did you hear about us?
          <select id="referrer">
            <option value="">(select one)</option>
            <option value="1">freeCodeCamp News</option>
            <option value="2">freeCodeCamp YouTube Channel</option>
            <option value="3">freeCodeCamp Forum</option>
            <option value="4">Other</option>
          </select>
        </label>
        <label for="bio">Provide a bio:
          <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
        </label>
      </fieldset>
      <label for="terms-and-conditions">
        <input id="terms-and-conditions" type="checkbox" required /> I accept the <a href="https://www.freecodecamp.org/news/terms-of-service/">terms and conditions</a>
      </label>
      --fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
