---
id: 60fab4a123ce4b04526b082b
title: Passaggio 28
challengeType: 0
dashedName: step-28
---

# --description--

Attualmente gli utenti possono inviare il modulo senza selezionare gli input radio. Anche se in precedenza hai usato l'attributo `required` per indicare che un input è obbligatorio, non funzionerà in questo caso perché aggiungere `required` a entrambi gli input trasmetterebbe informazioni errate agli utenti.

Per risolvere questo problema, puoi fornire il contesto di ciò che è necessario aggiungendo un elemento `legend` con il testo `Account type (required)` prima degli elementi `label` all’interno del secondo `fieldset`. Poi aggiungi l'attributo `checked` all’input `Personal` per assicurarti che il modulo venga inviato con i dati richiesti.

# --hints--

I tuoi elementi `input` dovrebbero rimanere `type` di `radio`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]')?.length, 2);
```

Dovresti aggiungere un elemento `legend` all’interno del secondo elemento `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset:nth-of-type(2) > legend')?.length, 1);
```

Dovresti aggiungere il testo `Account type (required)` all’elemento `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Dovresti dare un attributo `checked` all’input `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset>

        <label><input type="radio" name="account-type" /> Personal</label>
        <label><input type="radio" name="account-type" /> Business</label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
