---
id: 62cc5b1779e4d313466f73c5
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Annida un elemento `link` all'interno dell'elemento `head`. Dagli un attributo `rel` con valore `stylesheet` e un attributo `href` con valore `styles.css`.

# --hints--

Il tuo codice dovrebbe contenere un elemento `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Non dovresti modificare i tag `head` esistenti. Assicurati di non aver eliminato il tag di chiusura.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Dovresti avere un solo elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Il tuo elemento `link` dovrebbe essere un elemento vuoto, non dovrebbe avere un tag di chiusura `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Il tuo elemento `link` dovrebbe essere all'interno del tuo elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
