---
id: 6720b3191bbafd269bae0ae8
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Unda kipengele cha `div` chenye `id` ya `greeting-card` na `class` ya `card`.

Ndani ya kipengele cha `div`, ongeza `h1` yenye maandishi ya `Happy Birthday!`. Kisha ongeza kipengele cha aya chenye `class` kinachoitwa `message` na maandishi ya `Wishing you all the happiness and joy on your special day!`.

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert.exists(document.querySelector("div"));
```

Kipengele cha `div` kinapaswa kuwa na `id` ya `greeting-card`.

```js
assert.strictEqual(document.querySelector("div").id, "greeting-card");
```

Kipengele cha `div` kinapaswa kuwa na `class` ya `card`.

```js
assert.strictEqual(document.querySelector("#greeting-card").className, "card");
```

Unapaswa kuwa na kipengele cha `h1` ndani ya kipengele cha `div`.

```js
assert.exists(document.querySelector("div > h1"));
```

Kipengele cha `h1` kinapaswa kuwa na maandishi ya `Happy Birthday!`.

```js
assert.strictEqual(document.querySelector("div > h1").textContent.trim(), "Happy Birthday!");
```

Unapaswa kuwa na kipengele cha `p` ndani ya `div`.

```js
assert.exists(document.querySelector("div > p"));
```

Kipengele cha `p` kinapaswa kuwa na `class` ya `message`.

```js
assert.exists(document.querySelector("div > p.message"));
```

Kipengele cha `p` kinapaswa kuwa na maandishi ya `Wishing you all the happiness and joy on your special day!`.

```js
assert.strictEqual(document.querySelector("div > p.message").textContent.trim(), "Wishing you all the happiness and joy on your special day!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```

```css

```
