---
id: 6720b7f3155a9278d6b6a8b8
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa ni wakati wa kupamba kadi yako ya salamu.

Ongeza kichaguzi kwa kipengele cha `body`, kisha:

- badilisha `font-family` kuwa `Arial` ikifuatiwa na `sans-serif` ya jumla,
- toa nafasi ya ndani (padding) upande wote wa `40px`,
- weka sifa ya `text-align` kuwa `center`.

# --hints--

Unapaswa kuwa na kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Unapaswa kuweka `font-family` kuwa `"Arial", sans-serif` katika kichaguzi cha `body`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("font-family"), ["'Arial', sans-serif", "Arial, sans-serif", '"Arial", sans-serif']);
```

Unapaswa kuweka `padding` kuwa `40px` katika kichaguzi cha `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("padding"), "40px");
```

Unapaswa kuweka `text-align` kuwa `center` katika kichaguzi cha `body`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.getPropertyValue("text-align"), "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
