---
id: 6720b1f2680ec706ed534acf
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utazoea kufanya kazi na darasa la bandia na kipengele cha bandia tofauti kwa kubuni kadi ya salamu. Muundo wa msingi wa HTML umeandaliwa kwa ajili yako.

Anza warsha kwa kuunganisha faili la `styles.css`.

# --hints--

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `rel` yenye thamani `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele `href` yenye thamani `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.match(linkHrefValue, /^(\.\/)?styles\.css$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
  </body>
</html>
```

```css

```
