---
id: 672b4434dd6d6a1edc405419
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Ongeza vipengele viwili vya `a` ndani ya kipengele cha `.card-links`.

Kipengele cha kwanza kinapaswa kuwa na maandishi ya `Send Card`, `class` ya `send-link` na `href` ya `#send`.

Kipengele cha pili kinapaswa kuwa na maandishi ya `Share on Social Media`, `class` ya `share-link` na `href` ya `#share`.

# --hints--

Kipengele cha kwanza cha `a` ndani ya `.card-links` kinapaswa kuwa na maandishi ya `Send Card`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].innerText, "Send Card");
```

Kipengele cha kwanza cha `a` kinapaswa kuwa na `class` yenye thamani ya `send-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[0].className, "send-link");
```

Kipengele cha kwanza cha `a` kinapaswa kuwa na `href` yenye thamani ya `#send`.

```js
assert.match(document.querySelectorAll(".card-links > a")[0].href, /#send$/);
```

Kipengele cha pili cha `a` ndani ya `.card-links` kinapaswa kuwa na maandishi ya `Share on Social Media`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].innerText, "Share on Social Media");
```

Kipengele cha pili cha `a` kinapaswa kuwa na `class` yenye thamani ya `share-link`.

```js
assert.strictEqual(document.querySelectorAll(".card-links > a")[1].className, "share-link");
```

Kipengele cha pili cha `a` kinapaswa kuwa na `href` yenye thamani ya `#share`.

```js
assert.match(document.querySelectorAll(".card-links > a")[1].href, /#share$/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
--fcc-editable-region--
        
--fcc-editable-region--
      </div>
  	</div>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
