---
id: 672e44e811e610232fead333
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Ongeza `h2` ndani ya `#send` ambayo ina maandishi `Sending your card...`, kisha ongeza kipengele cha `p` chenye maandishi `Card successfully sent to your recipient!`.

# --hints--

Kipengele cha `#send` kinapaswa kuwa na kipengele cha `h2` ndani yake.

```js
assert.exists(document.querySelector("#send > h2"));
```

Kipengele cha `h2` kinapaswa kuwa na maandishi `Sending your card...`.

```js
assert.strictEqual(document.querySelector("#send > h2").innerText.trim(), "Sending your card...");
```

Kipengele cha pili ndani ya `#send` kinapaswa kuwa kipengele cha `p`.

```js
assert.strictEqual(document.querySelector("#send").children[1].tagName, "P");
```

Kipengele cha `p` kinapaswa kuwa na maandishi `Card successfully sent to your recipient!`.

```js
assert.strictEqual(document.querySelector("#send").children[1].innerText.trim(), "Card successfully sent to your recipient!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
    </div>
    <section id="send">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
    <section id="share"></section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
