---
id: 6733352ce2a5642827b4f7f0
title: Hatua 18
challengeType: 0
dashedName: step-18
---

# --description--

Ongeza kichaguzi cha `.card-links`, na weka sifa ya `margin-top` kuwa `20px`.

Unaweza kuongeza `display: flex` kuweka kipengele kitumie flexbox, utajifunza zaidi kuhusu flexbox baadaye katika kozi, unaweza kuichukulia hii kama muhtasari mdogo.

Ili kuweka nafasi sawa kati ya viungo viwili hivyo viwe na nafasi sawa pande zote, ongeza sifa ya `display` iwe `flex`, na `justify-content` iwe `space-around`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.card-links`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".card-links"));
```

Kichaguzi cha `.card-links` kinapaswa kuwa na sifa ya `margin-top` iwe `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("margin-top"), "20px");
```

Kichaguzi cha `.card-links` kinapaswa kuwa na sifa ya `display` iwe `flex`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("display"), "flex");
```

Kichaguzi cha `.card-links` kinapaswa kuwa na sifa ya `justify-content` iwe `space-around`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".card-links")?.getPropertyValue("justify-content"), "space-around");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
      <h2>Sharing your card...</h2>
      <p>Your card was shared on social media!</p>
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
  transition: transform 0.3s, background-color 0.3s ease
}

.card:hover {
  background-color: khaki;
  transform: scale(1.1);
}

h1::before {
  content: "🥳 ";
}

h1::after {
  content: " 🥳";
}

.message {
  font-size: 1.2em;
  margin-bottom: 20px;
}

--fcc-editable-region--

--fcc-editable-region--
```
