---
id: 67332db0f793bf08f2a18528
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Ni wakati wa kujaza `section` ya pili! 

Ongeza kipengele cha `h2` ndani ya kipengele cha `#share` chenye maandishi ya `Sharing your card...`, kisha ongeza kipengele cha `p` chenye maandishi ya `Your card was shared on social media!`.

# --hints--

Kipengele cha `#share` kinapaswa kuwa na kipengele cha `h2` ndani yake.

```js
assert.exists(document.querySelector("#share > h2"));
```

Kipengele cha `h2` kinapaswa kuwa na maandishi ya `Sharing your card...` ndani yake.

```js
assert.strictEqual(document.querySelector("#share > h2").innerText.trim(), "Sharing your card...");
```

Kipengele cha pili ndani ya `#share` kinapaswa kuwa kipengele cha `p`.

```js
assert.strictEqual(document.querySelector("#share").children[1].tagName, "P");
```

Kipengele cha `p` kinapaswa kuwa na maandishi ya `Your card was shared on social media!`.

```js
assert.strictEqual(document.querySelector("#share").children[1].innerText.trim(), "Your card was shared on social media!");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Greeting Card</title>
    <link href="styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="card" id="greeting-card">
      <h1>Happy Birthday!</h1>
      <p class="message">
        Wishing you all the happiness and joy on your special day!
      </p>
      <div class="card-links">
        <a href="#send" class="send-link">Send Card</a>
        <a href="#share" class="share-link">Share on Social Media</a>
      </div>
  	</div>
    <section id="send">
      <h2>Sending your card...</h2>
      <p>Card successfully sent to your recipient!</p>
    </section>
    <section id="share">
  --fcc-editable-region--
      
  --fcc-editable-region--
    </section>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  padding: 40px;
  text-align: center;
  background-color: brown;
}

.card {
  background-color: white;
  max-width: 400px;
  padding: 40px;
  margin: 0 auto;
  border-radius: 10px;
  box-shadow: 0 4px 8px gray;
}

```
