---
id: 5a24c314108439a4d4036161
title: Aprende sobre las etiquetas de autocierre en JSX
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

Hasta ahora, has visto cómo JSX difiere de HTML de una manera clave con el uso de `className` vs. `class` para definir clases HTML.

Otra forma importante en la que JSX difiere de HTML está en la idea de la etiqueta de auto-cierre.

En HTML, casi todas las etiquetas tienen tanto una etiqueta de apertura como de cierre: `<div></div>`; la etiqueta de cierre siempre tiene una barra diagonal antes del nombre de la etiqueta que estás cerrando. Sin embargo, hay casos especiales en HTML llamados <dfn>elementos vacíos</dfn>, o elementos que no requieren tanto una etiqueta de apertura como de cierre antes de que pueda comenzar otro elemento.

Por ejemplo, la etiqueta de salto de línea puede escribirse como `<br>` o como `<br />`, pero nunca debe escribirse como `<br></br>`, ya que no contiene ningún contenido.

En JSX, las reglas son un poco diferentes. Cualquier elemento JSX puede escribirse con una etiqueta de autocierre, y cada elemento debe estar cerrado. La etiqueta de salto de línea, por ejemplo, siempre debe escribirse como `<br />` para que sea un JSX válido que pueda transpilarse. Un `<div>`, en cambio, puede escribirse como `<div />` o `<div></div>`. La diferencia es que en la primera versión de la sintaxis no hay forma de incluir nada dentro del `<div />`. Verás en desafíos posteriores que esta sintaxis es útil al renderizar componentes de React.

# --instructions--

Corrige los errores en el editor de código para que sea JSX válido y se transpile exitosamente. Asegúrate de no cambiar nada del contenido: sólo tienes que cerrar las etiquetas donde se necesite.

# --hints--

La constante `JSX` debe devolver un elemento `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

El `div` debe contener una etiqueta `br`.

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

El `div` debe contener una etiqueta `hr`.

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
