---
id: 5a24c314108439a4d4036185
title: Usar && para una condicional más concisa
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

Las sentencias `if/else` funcionaron en el último desafío, pero hay una forma más concisa de lograr el mismo resultado. Imagina que estás controlando varias condiciones en un componente y quieres que diferentes elementos se rendericen según cada una de estas condiciones. Si escribes muchas sentencias `else if` para devolver interfaces ligeramente diferentes, puedes repetir código, lo que deja espacio para errores. En lugar de eso, puedes usar el operador lógico `&&` para realizar lógica condicional de forma más concisa. Esto es posible porque quieres verificar si una condición es `true` y, si lo es, devolver un marcado. Aquí tienes un ejemplo:

```jsx
{condition && <p>markup</p>}
```

Si la `condition` es `true`, se devolverá el marcado. Si la condición es `false`, la operación devolverá inmediatamente `false` tras evaluar la `condition` y no devolverá nada. Puedes incluir estas sentencias directamente en tu JSX y encadenar múltiples condiciones escribiendo `&&` después de cada una. Esto te permite manejar lógica condicional más compleja en tu método `render()` sin repetir mucho código.

# --instructions--

Resuelve el ejemplo anterior de nuevo, de modo que el `h1` solo se renderice si `display` es `true`, pero usa el operador lógico `&&` en lugar de una estructura `if/else`.

# --hints--

`MyComponent` debe existir y renderizar.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

Cuando `display` se establece en `true`, un `div`, `button`, y `h1` debe renderizarse.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

Cuando `display` se establece en `false`, solo un `div` y un `button` debe renderizarse.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

El método render debe usar el operador lógico `&&` para verificar la condición de `this.state.display`.

```js
assert(code.includes('&&'));
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    };
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
