---
id: 60a3e3396c7b40068ad6996c
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

En lecciones anteriores sobre el modelo de caja de CSS, aprendiste que cada elemento HTML se trata como una caja con cuatro áreas.

Imagina que recibes una caja de tu minorista en línea favorito: el contenido es el artículo dentro de la caja, o en tu caso, un encabezado, párrafo o elemento de imagen.

Cambia el atributo `src` en el `<img>` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` a `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`.

# --hints--

El elemento `img` debe tener un `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png`

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-2.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png">
--fcc-editable-region--
  </body>
</html>
```

