---
id: 691f7bbf76229bb97c60827d
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتمرن على استخدام حلقات `for...of` من خلال بناء دالة تعد عدد مرات ظهور سلسلة نصية في مصفوفة من السلاسل النصية.

قبل كتابة تلك الدالة، أنشئ دالة أبسط تسجل كل حرف من سلسلة نصية في وحدة التحكم.

ابدأ بتعريف دالة فارغة تسمى `printCharacters` بمعلمة واحدة اسمها `str`.

# --hints--

يجب أن تنشئ دالة `printCharacters`.

```js
assert.isFunction(printCharacters);
```

يجب أن تحتوي دالة `printCharacters` على معلمة واحدة تسمى `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
