---
id: 6920d189e36b6b60a6330fee
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

الآن، داخل الحلقة، سجّل المتغير `char` في وحدة التحكم.

# --hints--

يجب أن تحتوي الحلقة على عبارة `console.log` مع `char` كمعلمتها.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

عند استدعاء الدالة `printCharacters` بالمعلمة `'test'`، يجب استدعاء عبارة `console.log()` داخل الدالة أربع مرات بالمعلمات التالية، بالترتيب: `'t'`، `'e'`، `'s'`، و`'t'`. بدلاً من ذلك، تم استدعاؤها بالمعلمات التالية: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
