---
id: 6920c9cda1c7f056f75e4b1a
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

حان الوقت الآن لكتابة أول حلقة `for...of` لك. في الدروس السابقة، تعلمت أنه يمكنك استخدام حلقة `for...of` للتكرار على القيم من قابل للتكرار (مثل السلاسل النصية أو المصفوفات).

إليك مثال على حلقة `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

أضف حلقة `for...of` مع كتلة شفرة فارغة داخل دالتك. يجب أن تتكرر على كل حرف في المعلمة `str`.

# --hints--

يجب أن تحتوي دالة `printCharacters` على حلقة `for...of` داخل كتلة الشفرة الخاصة بها.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

يجب أن تتكرر حلقة `for...of` على المعلمة `str` في دالتك.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
