---
id: 6921dae10700e300f26994e3
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

لعرض ناتج دالة `getMatchedWordCount` في حالتها الحالية، أضف عبارة `console.log` أسفلها. داخل `console.log()`، استدعِ الدالة مع المعلمة `["I", "really", "really", "really", "like", "to", "code"]` للمعلمة `sentence` و`"really"` للمعلمة `match`.

# --hints--

يجب استدعاء الدالة `getMatchedWordCount` داخل عبارة `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

يجب استدعاء الدالة `getMatchedWordCount` مع المعلمة `["I", "really", "really", "really", "like", "to", "code"]` للمعلمة `sentence` و`"really"` للمعلمة `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
