---
id: 6921e10ae2681205b516fbe6
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

حان الوقت الآن لإضافة المنطق إلى الحلقة التي تزيد قيمة `count` بشكل مناسب.

داخل الحلقة، استخدم عبارة شرطية لزيادة `count` بمقدار `1` إذا كانت قيمة المتغير `word` تساوي قيمة المتغير `match`. وإلا، اترك قيمة `count` كما هي.

# --hints--

يجب أن تُعيد الاستدعاءات `getMatchedWordCount([], "bar")`، و`getMatchedWordCount(["foo"], "bar")`، و`getMatchedWordCount(["foo", "bar", "foo"], "bar")`، و`getMatchedWordCount(["foo", "bar", "foo"], "foo")` بالترتيب قيمة من النوع `"number"`؛ لكن بدلاً من ذلك، أعادت قيمًا من النوع `--fcc-actual--`.

```js
const obj = { method: (arg1, arg2) => typeof getMatchedWordCount(arg1, arg2) };
const spy = __helpers.spyOn(obj, 'method');
obj.method([], "bar");
obj.method(["foo"], "bar");
obj.method(["foo", "bar", "foo"], "bar");
obj.method(["foo", "bar", "foo"], "foo");
assert.deepEqual(spy.returns, ["number", "number", "number", "number"]);
spy.restore();
```

يجب أن تُعيد `getMatchedWordCount([], "bar")` القيمة `0`، لكنها تُعيد `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount([], "bar"));
assert.strictEqual(resultString, "0");
```

يجب أن تُعيد `getMatchedWordCount(["foo"], "bar")` القيمة `0`، لكنها تُعيد `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo"], "bar"));
assert.equal(resultString, "0");
```

يجب أن تُعيد `getMatchedWordCount(["foo", "bar", "foo"], "bar")` القيمة `1`، لكنها تُعيد `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "bar"));
assert.equal(resultString, "1");
```

يجب أن تُعيد `getMatchedWordCount(["foo", "bar", "foo"], "foo")` القيمة `2`، لكنها تُعيد `--fcc-actual--`.

```js
const resultString = String(getMatchedWordCount(["foo", "bar", "foo"], "foo"));
assert.equal(resultString, "2");
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
