---
id: 6921d90b5ce50185505d2004
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

بما أن الدالة `getMatchedWordCount` تحتاج إلى إرجاع عدد رقمي، ستبدأ بإعداد متغير عداد.

أنشئ متغيرًا باسم `count` باستخدام الكلمة المفتاحية `let`، وابدأ بقيمة `0`، وأضف جملة `return` تُرجع المتغير `count`.

# --hints--

يجب إعلان المتغير `count` داخل الدالة `getMatchedWordCount`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getMatchedWordCount?.variables.count);
```

للسماح بزيادة المتغير `count` لاحقًا، أعلن عنه باستخدام الكلمة المفتاحية `let`.

```js
const explorer = await __helpers.Explorer(code);
const countVar = explorer.allFunctions.getMatchedWordCount?.variables.count;
assert.isTrue(countVar?.toString().startsWith('let'));
```

يجب أن تحتوي دالتك على تعيين القيمة `0` للمتغير `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.variables.count?.value.matches('0'));
```

يجب أن تحتوي دالتك على جملة `return` تُرجع المتغير `count`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getMatchedWordCount?.hasReturn('count'));
```

يجب أن تُرجع دالة `getMatchedWordCount` القيمة `0`.

```js
assert.strictEqual(getMatchedWordCount(["foo", "bar"], "foo"), 0);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
