---
id: 6921d688b8d58581d34459b9
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

في الخطوات القليلة التالية، ستبني دالة تحسب عدد مرات ظهور سلسلة نصية في مصفوفة من السلاسل النصية.

لتبدأ، عرّف دالة فارغة باسم `getMatchedWordCount` تحتوي على المعلمتين `sentence` و `match` بهذا الترتيب.

# --hints--

يجب أن تنشئ دالة `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

يجب أن تحتوي دالة `getMatchedWordCount` على معلمتين: `sentence` و `match`، بهذا الترتيب.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
