---
id: 6921de8d2a111d0341cff702
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

رائع، دالتك تُرجع `0`، لكنها لا تحسب أي شيء بعد.

لإصلاح ذلك، داخل `getMatchedWordCount`، أنشئ حلقة `for...of` مع كتلة شفرة فارغة تتكرر على كل كلمة في `sentence`.

# --hints--

يجب أن تحتوي دالة `getMatchedWordCount` على حلقة `for...of` داخل كتلة الشفرة الخاصة بها.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

يجب أن تتكرر حلقة `for...of` على المعلمة `sentence` في دالتك.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;

--fcc-editable-region--
  
--fcc-editable-region--

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
