---
id: 69220a03729c43082ad3e9f6
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

لقد أنهيت العمل على دالة `getMatchedWordCount` الخاصة بك!

الآن ستختبرها ببيانات مختلفة لترى كيف تتصرف.

أضف عبارة `console.log` جديدة أسفل العبارة الحالية التي تستدعي `getMatchedWordCount` بهذه المعلمات:

- sentence: `["Do", "not", "fear", "the", "dandy", "lion"]`
- match: `"dandy"`

تهانينا! لقد أكملت هذه الورشة.

# --hints--

يجب أن يكون لديك عبارة `console.log` ثانية تستدعي دالة `getMatchedWordCount` بداخلها.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.exists(normalizedConsoleLogCode);
```

في عبارة `console.log` الثانية، استدع دالة `getMatchedWordCount` مع `["Do", "not", "fear", "the", "dandy", "lion"]` كمعلمة `sentence` و`"dandy"` كمعلمة `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/g)[1];
assert.match(normalizedConsoleLogCode, /\[('|"|`)Do\1,('|"|`)not\2,('|"|`)fear\3,('|"|`)the\4,('|"|`)dandy\5,('|"|`)lion\6\],('|"|`)dandy\7/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
    if (word === match) {
      count++;
    }
    console.log(`Checking "${word}" against "${match}" | Running count: ${count}`);
  }
  
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);

console.log(
  getMatchedWordCount(
    ["Do", "not", "fear", "the", "dandy", "lion"],
    "dandy"
  )
);
```
