---
id: 6978bd465213832659f3338b
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

الآن، داخل الحلقة، سجّل في وحدة التحكم سلسلة نصية قالبية كما يلي: `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

يجب أن يكون لديك بالضبط بيان `console.log` واحد داخل الحلقة في دالتك `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

يجب أن يستخدم بيان `console.log` داخل الحلقة في دالتك `getMatchedWordCount` سلسلة نصية قالبية تحتوي على عناصر نائب التالية: `${word}`، `${match}`، و`${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

عندما تُستدعى الدالة `getMatchedWordCount` مع `(['foo'], 'bar')`، يجب أن يُستدعى بيان `console.log()` داخل الدالة مرة واحدة مع السلسلة النصية `Checking "foo" against "bar" | Running count: 0`. بدلاً من ذلك، تم استدعاؤه مع كل سلسلة في القائمة المفصولة بفواصل التالية: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
