---
id: 65823634a308cb4d6bcc5fa6
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Лямбда-функції корисно поєднувати з функцією `map()`, яка виконує задану функцію для кожного елемента в колекції об’єктів, наприклад, у списку:

```py
map(lambda x: x * 2, [1, 2, 3])
```

Функція для виконання передається як перший аргумент, а ітерований об’єкт — як другий аргумент.

Результатом наведеного вище прикладу буде `[2, 4, 6]`, де кожен елемент у списку, переданому `map()`, був подвоєний дією лямбда-функції.

Змініть свій виклик `print()`, щоб вивести результат виклику `map()` з `test` як першим аргументом і `[2, 3, 5, 8]` як другим аргументом. Поки що ви не зможете побачити читабельний вихідний результат.

# --hints--

У вашому коді має бути `print(map(test, [2, 3, 5, 8]))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*map\s*\(\s*test\s*,\s*\[\s*2\s*,\s*3\s*,\s*5\s*,\s*8\s*\]\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(test(3))
--fcc-editable-region--

expenses = []
```
