---
id: 6582267ef39f5047411c1743
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Далі визначте функцію з іменем `print_expenses`, яка приймає один параметр `expenses`. Цю функцію пізніше використовуватимуть для відображення кожного витрату у вашому списку.

Заповніть тіло нової функції інструкцією `pass`.

# --hints--

Ви повинні визначити нову функцію з іменем `print_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(print_expenses)
  `))
})
```

Ваша функція `print_expenses` має приймати один параметр `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(print_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
