---
id: 65822308ded1ba4632f66aa3
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Словник — це ще один вбудований тип даних у Python. Словник — це колекція даних у формі пар *ключ*-*значення*. Словники визначаються за допомогою фігурних дужок (`{}`) і містять пари ключ-значення, розділені комами. Кожен ключ супроводжується двокрапкою (`:`) і значенням:

```py
{'amount': 50.0, 'category': 'Food'}
```

У наведеному вище прикладі `'amount'` і `'category'` — це ключі, а `50.0` і `'Food'` — відповідні їм значення.

Створіть словник із ключем `'amount'` і значенням параметра `amount` та передайте свій новий словник у виклик `.append()`.

# --hints--

Ви повинні передати словник як аргумент методу `.append()`.

```js
({ test: () => assert(runPython(`
import ast
node = _Node(_code).find_function("add_expense").find_calls("append")[0].find_call_args()
len(node) == 1 and isinstance(node[0].tree, ast.Dict)
`)) })
```

Ви повинні передати `{'amount': amount}` як аргумент методу `.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append({'amount': amount})")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    expenses.append()
--fcc-editable-region--

expenses = []
```
