---
id: 65822773f4c60a479ec89f66
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Усередині функції `print_expenses` створіть цикл `for`, який ітерує кожен елемент у списку `expenses`. Використовуйте `expense` як змінну циклу та перемістіть `pass` всередину тіла циклу.

# --hints--

Ви повинні створити цикл `for`, який ітерує `expenses` у вашій функції `print_expenses`. Пам’ятайте використовувати `pass` у тілі циклу.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_iter().is_equivalent("expenses")`)) })
```

Ви повинні використовувати `expense` як змінну циклу.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_for_vars().is_equivalent("expense")`)) })
```

`pass` має бути лише всередині тіла циклу.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    pass
--fcc-editable-region--

expenses = []
```
