---
id: 65823040bdb32949d4c1a96f
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Вам потрібна функція для обчислення загальної суми витрат.

Визначте функцію з іменем `total_expenses`, яка приймає один параметр `expenses`. Поки що заповніть тіло функції інструкцією `pass`.

# --hints--

Ви повинні визначити функцію з іменем `total_expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(total_expenses)
  `))
})
```

Ваша функція `total_expenses` має приймати один параметр `expenses`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(total_expenses))
    sig == '(expenses)'
  `))
})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
