---
id: 658212ba98182f3e855e85f9
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

У цьому проєкті ви створите простий, але функціональний трекер витрат на Python.

Почніть з оголошення функції з іменем `add_expense`, яка приймає три параметри: `expenses`, `amount` і `category`. Використайте ключове слово `pass` для заповнення тіла функції.

# --hints--

Ви повинні використати ключове слово `def` для оголошення функції з іменем `add_expense`.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(add_expense)
  `))
})
```

Ваша функція `add_expense` має приймати три параметри: `expenses`, `amount` і `category`.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(add_expense))
    sig == '(expenses, amount, category)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
