---
id: 65821c1c40ccac44adf38e6b
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Створіть порожній список з іменем `expenses`. Ви будете використовувати його для збереження кожного зі своїх витрат.

# --hints--

Ви повинні оголосити нову змінну з іменем `expenses`.

```js
({ test: () => assert(__userGlobals.has("expenses")) })
```

Ви повинні призначити порожній список змінній `expenses`.

```js
({ test: () => assert.equal(__userGlobals.get("expenses").length, 0) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    pass
    
--fcc-editable-region--

--fcc-editable-region--
```
