---
id: 65822934c8a0454803e3ef12
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Далі ви маєте відобразити деталі для кожної витрати.

Всередині циклу `for` замініть `pass` на виклик `print()` і передайте йому наступний f-рядок: `f'Amount: {expense}, Category: {expense}'`.

# --hints--

Ви маєте вивести `f'Amount: {expense}, Category: {expense}'` у вашому циклі `for`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].has_call("print(f'Amount: {expense}, Category: {expense}')")`)) })
```

У тілі циклу не повинно бути `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("print_expenses").find_for_loops()[0].find_bodies()[0].has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        pass
--fcc-editable-region--

expenses = []
```
