---
id: 658246c28575d653d1f89b59
title: Крок 33
challengeType: 20
dashedName: step-33
---

# --description--

Цикл `while` — це інший тип циклу, який виконує частину коду доти, доки вказана умова є `True`. Цикл завершується, коли умова стає `False`:

```py
while condition:
    <code>
```

Під списком `expenses` створіть цикл `while`. Використайте `True` як умову і виведіть рядок `'\nExpense Tracker'` у тілі циклу, щоб показати назву програми.

# --hints--

Ви повинні створити цикл `while`, використовуючи `True` як умову у вашій функції `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_conditions()[0].is_equivalent("True")`)) })
```

Ви повинні вивести `'\nExpense Tracker'` у вашому циклі `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_calls("print")[0].is_equivalent("print('\\\\nExpense Tracker')")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--
def main():
    expenses = []
    
--fcc-editable-region--
```
