---
id: 65821fcc010c3245718f2a06
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Параметр `expenses` вашої функції `add_expense` буде списком витрат. Ви хочете мати можливість додавати елементи в кінець свого списку. Для цього ви використаєте метод списку `.append()`.

Додайте виклик методу `.append()` для списку `expenses`. Поки що не передавайте жодних аргументів методу `.append()`.

# --hints--

У вашій функції має бути `expenses.append()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("add_expense").has_stmt("expenses.append()")`)) })
```

У вашій функції не має бути `pass`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("add_expense").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def add_expense(expenses, amount, category):
    pass
--fcc-editable-region--

expenses = []
```
