---
id: 6582507654b3ed5712341382
title: Крок 38
challengeType: 20
dashedName: step-38
---

# --description--

Ви збираєтеся використати умовні інструкції, щоб перевірити вибір користувача. Якщо вибір дорівнює `'1'`, це означає, що користувач хоче додати витрату.

Ще в циклі `while`, під змінною `choice`, напишіть інструкцію `if`, щоб перевірити, чи дорівнює `choice` рядку `'1'`. Якщо це правда, це буде початком додавання нової витрати.

У тілі інструкції `if` оголосіть змінну `amount` і призначте їй виклик пустого `input()`.

# --hints--

Ви повинні створити інструкцію `if`, яка перевіряє, чи дорівнює `choice` `'1'` у вашому циклі `while`.

```js
({ test: () => assert(runPython(`
cond = _Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
cond.is_equivalent("choice == '1'") or cond.is_equivalent("'1' == choice")
`)) })
```

Ви повинні оголосити змінну з ім’ям `amount` у тілі вашої інструкції `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("amount")
`)) })
```

Ви повинні призначити виклик пустого `input()` вашій змінній `amount`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("amount").is_equivalent("amount = input()")
`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    
--fcc-editable-region--
def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')

--fcc-editable-region--
```
