---
id: 65824111a09164518320088d
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Функція `filter()` дозволяє вибирати елементи з ітерованого об’єкта, такого як список, на основі вихідних даних функції:

```py
filter(my_function, my_list)
```

`filter()` приймає функцію як перший аргумент і ітерований об’єкт як другий аргумент. Вона повертає ітератор — спеціальний об’єкт, який дає змогу ітерувати над елементами колекції, наприклад списку.

Результатом прикладу вище є ітератор, що містить елементи `my_list`, для яких `my_function` повертає `True`.

Усередині функції `filter_expenses_by_category` викличте `filter()`, передавши функцію `lambda`, яку ви написали на попередньому кроці, як перший аргумент, а список `expenses` — як другий аргумент.

# --hints--

Ви повинні викликати `filter()` всередині функції `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")) == 1`)) })
```

Ви повинні передати `lambda expense: expense['category'] == category` як перший аргумент виклику `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[0].is_equivalent("lambda expense: expense['category'] == category")`)) })
```

Ви повинні передати `expenses` як другий аргумент виклику `filter()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").find_calls("filter")[0].find_call_args()[1].is_equivalent("expenses")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    lambda expense: expense['category'] == category
--fcc-editable-region--

expenses = []
```
