---
id: 65824561f3478e5371a33ae5
title: Крок 32
challengeType: 20
dashedName: step-32
---

# --description--

Наступний крок — визначити головну функцію, яка буде точкою входу в інтерактивну програму для відстеження витрат.

Визначте функцію з іменем `main` без параметрів. Заповніть тіло функції списком `expenses`, який ви створили на початку цього проєкту. Ви будете використовувати цей список для збереження записів про витрати.

# --hints--

Ви повинні визначити функцію з іменем `main()` без параметрів.

```js
({ test: () => assert(runPython(`
    import inspect    
    inspect.isfunction(main)
    sig = str(inspect.signature(main))
    sig == '()'
  `))
})
```

Ви повинні перемістити список `expenses` всередину тіла функції `main()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_body().is_equivalent("expenses = []")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
--fcc-editable-region--

expenses = []
--fcc-editable-region--
```
