---
id: 6582324b5d37fe4b622bc3c0
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Lambda-функції — це короткі анонімні функції в Python, ідеальні для простих одноразових завдань. Вони оголошуються за допомогою ключового слова `lambda` і мають такий синтаксис:

```py
lambda x: expr
```

У наведеному прикладі `x` представляє параметр, який використовується у виразі `expr`, і він поводиться так само, як будь-який параметр у традиційній функції. `expr` — це вираз, який обчислюється і повертається під час виклику lambda-функції.

Створіть змінну з іменем `test` і призначте їй lambda-функцію, яка приймає параметр `x` і повертає `x * 2`.

# --hints--

Ви повинні оголосити змінну з іменем `test`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("test")`))})
```

Ви повинні призначити `lambda x: x * 2` вашій змінній `test`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("test").is_equivalent("test = lambda x: x * 2")`))})
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--

--fcc-editable-region--

expenses = []
```
