---
id: 65822bd82d708c4895080c35
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Важливо знати, що в Python не можна використовувати той самий тип лапок для визначення рядка і всередині нього. Наприклад, рядок `'I'm a string!'` є недійсним. Щоб використати одинарну лапку всередині такого рядка, потрібно або:

- Екранувати лапку, додавши перед нею зворотний слеш: `'I\'m a string!'`
- Або визначити рядок за допомогою подвійних лапок: `"I'm a string!"` (рекомендується).

Ви можете отримати доступ до значень у словнику через його ключі. Для цього потрібно використати дужкову нотацію, включивши ключ у квадратні дужки:

```py
my_dict = {'amount': 50.0, 'category': 'Food'}
my_dict['amount'] # 50.0
```

Зараз ви інтерполюєте словник `expense` у своєму f-рядку. Змініть вираз у f-рядку, щоб отримати значення ключа `'amount'` і ключа `'category'` у словнику `expense`.

# --hints--

Ви повинні передати `f'Amount: {expense["amount"]}, Category: {expense["category"]}'` у свій виклик `print()`. Пам’ятайте, що всередині f-рядка з одинарними лапками слід використовувати подвійні лапки і навпаки.

```js
({ test: () => assert.match(code, /^\s+print\s*\(\s*f("|')Amount: \{\s*expense\s*\[\s*(?=[^\1])("|')amount\2\s*\]\s*\}, Category: \{\s*expense\s*\[\s*(?=[^\1])("|')category\3\s*\]\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
--fcc-editable-region--
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense}, Category: {expense}')
--fcc-editable-region--

expenses = []
```
