---
id: 6582575b8089f85b8b92d7c8
title: Крок 42
challengeType: 20
dashedName: step-42
---

# --description--

Після того, як ви отримаєте деталі витрат, потрібно викликати функцію `add_expense`, щоб додати нову витрату до списку `expenses`.

Після отримання `amount` і `category` за допомогою `input()`, викличте функцію `add_expense`, передаючи три аргументи: `expenses`, `amount` і `category`.

- `expenses` — це порожній список, створений у головній функції раніше в цьому проєкті.
- `amount` — це сума витрати.
- `category` — це категорія витрати.

# --hints--

Ви повинні викликати функцію `add_expense()` у вашій інструкції `if`.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_calls("add_expense")) == 1`)) })
```

Ви повинні передати три аргументи у виклик `add_expense()`: `expenses`, `amount` і `category`. Порядок має значення.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("add_expense(expenses, amount, category)")`)) })
```

Ваш виклик `add_expense()` має йти після присвоєння `amount` і `category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("main").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].is_equivalent("amount = float(input('Enter amount: '))\\ncategory = input('Enter category: ')\\nadd_expense(expenses, amount, category)")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
def filter_expenses_by_category(expenses, category):
    return filter(lambda expense: expense['category'] == category, expenses)
    

def main():
    expenses = []
    while True:
        print('\nExpense Tracker')
        print('1. Add an expense')
        print('2. List all expenses')
        print('3. Show total expenses')
        print('4. Filter expenses by category')
        print('5. Exit')
        
        choice = input('Enter your choice: ')
--fcc-editable-region--
        if choice == '1':
            amount = float(input('Enter amount: '))
            category = input('Enter category: ')
            
--fcc-editable-region--
```
