---
id: 6582352cbdbcde4d0c4f7b0b
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Щоб викликати lambda-функцію, ви можете використати звичайний синтаксис функції з парою круглих дужок після імені змінної.

Викличте вашу `test` lambda-функцію та передайте `3` як аргумент. Потім виведіть результат.

# --hints--

У вашому коді має бути `print(test(3))`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*test\s*\(\s*3\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2

--fcc-editable-region--

expenses = []
```
