---
id: 65823ff0d4b991510fade1a8
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

У функції `filter_expenses_by_category` замініть `pass` на функцію `lambda`. Використайте `expense` як параметр і оцініть порівняння між значенням ключа `'category'` словника `expense` та `category` за допомогою оператора рівності.

# --hints--

Ви маєте створити функцію `lambda`, яка використовує параметр з іменем `expense` і оцінює вираз `expense['category'] == category` всередині функції `filter_expenses_by_category`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_stmt("lambda expense: expense['category'] == category")`)) })
```

У функції `pass` не повинно бути `filter_expenses_by_category`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_function("filter_expenses_by_category").has_pass()`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')
    
def total_expenses(expenses):
    return sum(map(lambda expense: expense['amount'], expenses))
    
--fcc-editable-region--
def filter_expenses_by_category(expenses, category):
    pass
--fcc-editable-region--

expenses = []
```
