---
id: 66314f2add0a8b57da5d6d91
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Ви повинні побачити щось на кшталт `<map object at 0xd273a8>`, виведене в консоль, що є рядковим представленням об’єкта map, який повертає `map()`.

Щоб отримати читабельний вихід, потрібно перетворити об’єкт map на список. Зробіть це, передавши виклик `map()` як аргумент функції `list()`.

# --hints--

Ви повинні передати виклик `map()` як аргумент функції `list()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(list(map(test, [2,3,5,8])))")`)) })
```

# --seed--

## --seed-contents--

```py
def add_expense(expenses, amount, category):
    expenses.append({'amount': amount, 'category': category})
    
def print_expenses(expenses):
    for expense in expenses:
        print(f'Amount: {expense["amount"]}, Category: {expense["category"]}')

def total_expenses(expenses):
    pass
    
--fcc-editable-region--
test = lambda x: x * 2
print(map(test, [2, 3, 5, 8]))
--fcc-editable-region--

expenses = []
```
