---
id: 66694fc4bba24f33ca01fa5b
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У Python є кілька методів для списків. Наприклад, методи для додавання або видалення елементів списку.

Ви можете додати елемент у кінець списку, використовуючи метод `append()`. Наприклад:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` тепер є `[4, 5, 6, 7]`.

Спробуйте використати метод `append()`, щоб додати `3` до `my_list`. Потім виведіть список.

# --hints--

Останній елемент `my_list` має бути `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` має мати довжину 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Ви не повинні змінювати оголошення списку.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Ви повинні вивести список.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Ви повинні вивести список після використання `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
