---
id: 66694fc4bba24f33ca01fa5b
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Python має пригорщ методів списку. Такі як методи додавання або видалення елементів списку.

Ви можете додати елемент у кінець списку за допомогою методу `append()`. Наприклад:

```py
example_list = [4, 5, 6]

example_list.append(7)
```

`example_list` тепер `[4, 5, 6, 7]`.

Спробуйте скористатися методом `append()`, щоб додати `3` до `my_list`. Потім, надрукуйте список.

# --hints--

Останнім елементом `my_list` має бути `3`.

```js
({
    test: () => runPython(`
        assert my_list[2] == 3, "my_list[2] is not 3"
    `)
})
```

`my_list` повинен мати довжину 3.

```js
({
    test: () => runPython(`
        assert len(my_list) == 3, "my_list doesn't have three elements"
    `)
})
```

Не змінюйте останнє оголошення.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Надрукуйте список.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls("print")
    assert any(call.is_equivalent('print(my_list)') for call in calls)
    `)
})
```

Надрукуйте список після використання `append()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.append(3)', 'print(my_list)')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]
--fcc-editable-region--
```
