---
id: 6669539c1379793f9cb8917c
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Списки в Python є змінними, що означає, що значення елементів списку можна змінювати. Ви можете змінити значення елемента, використовуючи дужкову нотацію.

```py
example_list = [4, 5, 6, 7]
example_list[1] = 'oh'
```

Це зробить так, що `example_list` матиме значення `[4, 'oh', 6, 7]`.

Змініть перший елемент `my_list` на `0`, а потім виведіть список, щоб перевірити значення.

# --hints--

Не змінюйте оголошення списку.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

Перший елемент `my_list` має бути 0.

```js
({
    test: () => runPython(`
        assert my_list[0] == 0
    `)
})
```

Ви повинні змінити перший елемент, використовуючи `my_list[0] = 0`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_stmt("my_list[0] = 0")
    `)
})
```

Ви повинні вивести список.

```js
({
    test: () => runPython(`
        calls = _Node(_code).find_calls('print')
        assert any(call.is_equivalent("print(my_list)") for call in calls)
    `)
})
```

Виведіть список після зміни першого елемента.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered("my_list[0] = 0", "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

--fcc-editable-region--
```
