---
id: 666951d2e4b27e3a3c72ca5f
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ви також можете отримати доступ до одного елемента, щоб отримати його значення.

Списки індексуються з нуля, як і рядки. Це означає, що перший елемент має індекс `0`, другий елемент має індекс `1` і так далі.

Щоб отримати доступ до елемента, ви використовуєте дужки. Наприклад, `example_list[1]` отримує доступ до елемента за індексом `1`, другого елемента `example_list`.

Надрукуйте перший елемент `my_list`.

# --hints--

Надрукуйте перший елемент списку за допомогою `print(my_list[0])`.

```js
({
    test: () => runPython(`
calls =  _Node(_code).find_calls("print")
assert any(call.is_equivalent('print(my_list[0])') for call in calls), "print of my_list[0] not found"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)
--fcc-editable-region--
```
