---
id: 66694eaaeb7d253049fa7883
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Щоб додати елементи до списку, ви можете написати їх між квадратними дужками, розділяючи значення комами.

Додайте числа `1` і `2` всередину списку.

# --hints--

Перший елемент `my_list` має бути `1`.

```js
({
    test: () => runPython(`
        assert my_list[0] == 1, "my_list[0] is not 1"
    `)
})
```

Другий елемент `my_list` має бути `2`.

```js
({
    test: () => runPython(`
        assert my_list[1] == 2, "my_list[1] is not 2"
    `)
})
```

`my_list` має містити два елементи всередині.

```js
({
    test: () => runPython(`
        assert len(my_list) == 2, "my_list doesn't have two elements"
    `)
})
```

Ви маєте призначити `[1, 2]` безпосередньо `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent("my_list = [1, 2]")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = []
--fcc-editable-region--
```
