---
id: 66694bfa435cef241f6f9a68
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Перед початком проєкту ви дізнаєтеся про списки.

На відміну від чисел і рядків, <dfn>список</dfn> — це тип даних, який працює як контейнер для інших значень:

```py
example_list = [1, 2, 3, 4]
empty_list = []
```

Список характеризується квадратними дужками навколо всіх значень і комою між значеннями, наприклад: `["A", "happy", "list"]`. Якщо список не містить жодних значень, тоді це порожній список: `[]`.

Список може містити різні типи даних: `[1, "Up", ["Down", "Twice"]]`. Це включає всі можливі типи даних. Він також може містити інший список!

Створіть змінну з назвою `my_list` і призначте їй порожній список.

# --hints--

У вас має бути змінна з назвою `my_list`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('my_list'), "variable my_list not found"
    `)
})
```

Змінна `my_list` має містити значення порожнього списку.

```js
({
    test: () => runPython(`
        assert _Node(_code).find_variable('my_list').is_equivalent('my_list = []'), "variable my_list is not an empty list"
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
