---
id: 666ab82c874add3c95145213
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Метод `pop()` можна використати, щоб видалити елемент зі списку. За замовчуванням він видаляє останній елемент списку. Ви можете передати індекс як аргумент методу, і він видалить елемент за вказаним індексом.

```py
fruits_list = ["cherry", "lemon", "tomato", "apple", "orange"]

fruits_list.pop(2)

print(fruits_list) # ["cherry", "lemon", "apple", "orange"]
```

У цьому випадку `fruits_list.pop(2)` видаляє елемент за індексом `2` зі списку.

Використайте `pop()`, щоб видалити останній елемент із `my_list`, а потім виведіть `my_list`.

# --hints--

Ви повинні використати `pop()`.

```js
({
    test: () => runPython(`
    assert len(_Node(_code).find_calls('pop')) > 0
    `)
})
```

`my_list` має бути `[0, 1, 2]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2]
    `)
})
```

Ви повинні вивести `my_list` після використання `pop()`.

```js
({
    test: () => runPython(`
    code = _Node(_code)
    assert(
        any(
            code.is_ordered(pop, 'print(my_list)') for pop in ['my_list.pop()', 'my_list.pop(3)']
        )
    )
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

my_list.insert(1, 1)
print(my_list)


--fcc-editable-region--
```
