---
id: 666aae6a5d30a71f1fd7749f
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Метод `insert` може додати елемент у будь-яку позицію списку. Перший аргумент — це позиція, на яку потрібно додати елемент, а другий аргумент — елемент, який потрібно додати. Наприклад, ось як додати новий елемент на третю позицію `example_list`:

```py
example_list = [4, 5, 6, 7]

example_list.insert(2, 5.5)

print(example_list) # [4, 5, 5.5, 6, 7]
```

Використовуючи `insert()`, додайте `1` до `my_list` у правильну позицію так, щоб список був відсортований за зростанням, а потім виведіть список.

# --hints--

Ви повинні додати `1` до `my_list` на індексі `1` за допомогою `insert()`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('my_list.insert(1, 1)')
    `)
})
```

`my_list` має містити 4 елементи.

```js
({
    test: () => runPython(`
    assert len(my_list) == 4
    `)
})
```

`my_list` має дорівнювати `[0, 1, 2, 3]`.

```js
({
    test: () => runPython(`
    assert my_list == [0, 1, 2, 3]
    `)
})
```

Ви повинні вивести список після додавання нового елемента.

```js
({
    test: () => runPython(`
    assert _Node(_code).is_ordered('my_list.insert(1, 1)', "print(my_list)")
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
my_list = [1, 2]

my_list.append(3)
print(my_list)

print(my_list[0])

my_list[0] = 0
print(my_list)

--fcc-editable-region--
```
