---
id: 66a9b8f14b963916a3baa732
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

Para adicionar uma legenda a uma `table`, você pode usar o elemento <dfn>caption</dfn> da tabela.

Aqui está um exemplo usando o elemento `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Dentro do elemento `table`, aninhe um elemento `caption` com o texto `Calculus Final Exam Grades`.

# --hints--

Você deve acrescentar uma tag de abertura para `caption`.

```js
assert.match(code, /<caption>/i);
```

Você deve acrescentar uma tag de fechamento para `caption`.

```js
assert.match(code, /<\/caption>/i);
```

O elemento `caption` deve estar dentro do elemento `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

O `caption` deve ter o texto `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText, 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    --fcc-editable-region--
    <table>

    </table>
    --fcc-editable-region--
  </body>
</html>
```
