---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Passo 9
challengeType: 0
dashedName: step-9
---

# --description--

Agora é hora de adicionar mais dois estudantes à tabela.

Following the same pattern as the previous step, add a second student table row. Use the following data for the table data elements:

- Sobrenome: `Doe`
- Nome: `Samantha`
- Nota: `92`

For the third student table row, use the following data for the table data elements:

- Sobrenome: `Rodriguez`
- Nome: `Marcus`
- Nota: `88`

# --hints--

You should have a second `tr` inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Você deve ter um elemento `td` com o texto `Doe` dentro do segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent, 'Doe');
```

Você deve ter um elemento `td` com o texto `Samantha` dentro do segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent, 'Samantha');
```

Você deve ter um elemento `td` com o texto `92` dentro do segundo elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent, '92');
```

You should have a third `tr` inside your `tbody` element.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Você deve ter um elemento `td` com o texto `Rodriguez` dentro do terceiro elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent, 'Rodriguez');
```

Você deve ter um elemento `td` com o texto `Marcus` dentro do terceiro elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent, 'Marcus');
```

Você deve ter um elemento `td` com o texto `88` dentro do terceiro elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent, '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
