---
id: 66a9cb26ec6bd41cf6c82bc5
title: Passo 10
challengeType: 0
dashedName: step-11
---

# --description--

A última seção a ser adicionada à tabela é a do elemento `table foot` (`<tfoot>`). O elemento `table foot` será usado para exibir a nota média de todos os alunos na tabela.

Adicione um elemento `tfoot` à tabela.

Dentro do elemento `tfoot`, adicione um elemento `tr`.

Dentro do elemento `tr`, adicione dois elementos `td`.

O primeiro elemento `td` deve conter o texto `Average Grade` (Média). O segundo elemento `td` deve conter o texto `78.8`.

# --hints--

Você deve ter uma tag de abertura de `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Você deve ter uma tag de fechamento de `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Você deve ter um elemento `tr` dentro do elemento `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

O primeiro elemento `td` deve conter o texto `Average Grade`. Certifique-se de que está dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').textContent, 'Average Grade');
```

O segundo elemento `td` deve conter o texto `78.8`. Certifique-se de que está dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)').textContent, '78.8');
```

O elemento `tr` deve ter dois elementos `td` dentro dele.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      --fcc-editable-region--

      --fcc-editable-region--
    </table>
  </body>
</html>
```
