---
id: 66a9c006e4f58819396605b3
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Agora que você completou a seção de cabeçalho, é hora de adicionar o `table body`. O `table body` representará todos os nomes dos estudantes e suas notas.

Adicione um elemento `table body` à tabela.

# --hints--

Você deve ter uma tag de abertura para `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Você deve ter uma tag de fechamento para `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

A tag de abertura de `tbody` deve vir antes da tag de fechamento de `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      --fcc-editable-region--

      --fcc-editable-region--
    </table>
  </body>
</html>
```
