---
id: 66a9c006e4f58819396605b3
title: Passo 7
challengeType: 0
dashedName: step-7
---

# --description--

Now that you have completed the head section, it is time to add the <dfn>table body</dfn>, `tbody`. The table body will represent all of the student names and their grades.

Add a table body element to your table.

# --hints--

Você deve ter uma tag de abertura para `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Você deve ter uma tag de fechamento para `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

A tag de abertura de `tbody` deve vir antes da tag de fechamento de `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
--fcc-editable-region--

--fcc-editable-region--
    </table>
  </body>
</html>
```
