---
id: 66a9c14e3b34c719e34bae20
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

To add student data to the table, you will need to use the table row and table data elements.

The <dfn>table data</dfn> element, `td`, is used to create a cell in the table.

Aqui está um exemplo de uso do elemento `td` para uma tabela de jogadores:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Inside your table body element, add a table row element, `tr`, with three table data elements, `td`.

The first table data element should contain the last name of `Davis`.

The second table data element should contain the first name of `Alex`.

The third table data element should contain the grade of `54`.

# --hints--

You should have a `tr` element inside your `tbody` element.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Você deve ter um elemento `td` com o texto `Davis` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent, 'Davis');
```

Você deve ter um elemento `td` com o texto `Alex` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent, 'Alex');
```

Você deve ter um elemento `td` com o texto `54` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent, '54');
```

Você deve ter todos os três elementos `td` dentro do elemento `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--

--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
