---
id: 66a98f42c7c06903e5f8dd07
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with HTML tables by building a final exam table for a group of students.

Para iniciar o projeto, adicione o `<!DOCTYPE html>` e um elemento `html` com um atributo `lang` de `en`.

Dentro do elemento `html`, adicione um elemento `head`.

# --hints--

Você deve ter a declaração `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter uma tag de abertura para `html` com o idioma definido para inglês.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve acrescentar uma tag de fechamento para `html`.

```js
assert.match(code, /<\/html>/i);
```

O `DOCTYPE` deve vir antes do elemento `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

Você deve acrescentar uma tag de abertura `head`.

```js
assert.match(code, /<head>/i);
```

Você deve ter uma tag de fechamento para `head`.

```js
assert.match(code, /<\/head>/i);
```

A tag de abertura de `head` deve vir antes da tag de fechamento de `head`.

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

O elemento `head` deve estar dentro do elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[.\n\s]*<head>[.\n\s]*<\/head>[.\n\s]*<\/html>/im);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
