---
id: 66a9c14e3b34c719e34bae20
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Para adicionar dados de estudantes à tabela, você precisa usar os elementos `table row` e <dfn>table data</dfn>.

O elemento `table data`, `<td>`, é usado para criar uma célula na tabela.

Aqui está um exemplo de uso do elemento `td` para uma tabela de jogadores:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Dentro do elemento `table body`, adicione um elemento `table row`, `<tr>`, com três elementos `table data`, `<td>`.

O primeiro elemento `table data` deve conter o sobrenome `Davis`.

O segundo elemento `table data` deve conter o primeiro nome `Alex`.

O terceiro elemento `table data` deve conter a nota `54`.

# --hints--

Você deve ter uma tag de abertura de `tr`.

```js
assert.match(code, /<tr>/i);
```

Você deve ter uma tag de fechamento de `tr`.

```js
assert.match(code, /<\/tr>/i);
```

O elemento `tr` deve estar dentro do elemento `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Você deve ter um elemento `td` com o texto `Davis` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td').textContent, 'Davis');
```

Você deve ter um elemento `td` com o texto `Alex` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1].textContent, 'Alex');
```

Você deve ter um elemento `td` com o texto `54` dentro do elemento `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2].textContent, '54');
```

Você deve ter todos os três elementos `td` dentro do elemento `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        --fcc-editable-region--

        --fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
