---
id: 69b83e35f19ba26ba1fa517a
title: 建置交通號誌序列器
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

在這個實驗中，你將模擬可設定的交通號誌週期並記錄異常狀況。

你將使用描述交通號誌階段的設定物件。每個設定物件具有以下屬性：

- `fault`：一個布林值旗標，當為 `true` 時會觸發提前終止。
- `phases`：一個階段物件的陣列。

`phases` 中的每個階段物件具有以下屬性：

- `color`：一個字串，代表燈號顏色（`"green"`、`"yellow"` 或 `"red"`）。
- `duration`：一個正整數，代表該階段持續的秒數。

你可以參考提供的 `config1`、`config2`、`config3` 和 `config4` 物件，作為可能的設定物件範例。

**注意**：請勿新增任何額外的 `console.log()` 陳述式，因為可能會導致測試失敗。

**目標**：完成以下使用者故事並通過所有測試以完成實驗。

# 使用者故事

1. 你應該有一個名為 `runSequence` 的函式，帶有兩個參數：`config` 和 `cycles`，其中 `config` 代表一個設定物件，`cycles` 代表序列最多可執行的次數。

1. 你應該使用 `for` 或 `while` 迴圈來實作 `runSequence(config, cycles)`，以迭代給定次數內的每個階段。

1. `runSequence` 函式應該：
    - 如果 `config.phases` 是空的，記錄 `No phases found` 並立即返回。
    - 如果 `config.fault` 設為 `true`，記錄 `Faulted phase!` 並提前停止模擬。
    - 如果 `duration <= 0`，記錄 `Invalid phase detected`。
    - 對每個有效階段，記錄 `Switching to [color] for [duration] s`。將 `[color]` 和 `[duration]` 替換為該階段物件的對應屬性。
    - 例如，`runSequence(config1, 1)` 應該記錄：

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. 你應該有一個名為 `generateTimeline` 的函式，帶有兩個參數：`config` 和 `cycles`。

1. `generateTimeline` 函式應該：
    - 在每個階段結束後，將累積經過時間記錄到陣列中，迭代時將每個階段的 `duration` 加到累計總和。
    - 即使 `config.fault` 是 `true` 或 `duration <= 0`，也要處理所有故障和無效階段，不做驗證。
    - 傳回累積時間戳的陣列。
    - 例如，`generateTimeline(config1, 1)` 應該傳回陣列 `[5, 7, 11]`。

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

你不應該更改 `config1`、`config2`、`config3` 和 `config4` 物件。

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

你應該有一個名為 `runSequence` 的函式。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.runSequence);
```

你的 `runSequence` 函式應該接受兩個參數，`config` 和 `cycles`。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.runSequence?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`runSequence(config1, 1)` 應該依序記錄 `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`。

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` 應該依序記錄 `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`、`Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`。

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` 應該依序記錄 `Switching to red for 3 s`、`Invalid phase detected`、`Switching to green for 6 s`。

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` 應該記錄 `Faulted phase!`。

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` 應該記錄 `No phases found`。

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

你應該有一個名為 `generateTimeline` 的函式。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.generateTimeline);
```

你的 `generateTimeline` 函式應該接受兩個參數，`config` 和 `cycles`。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.generateTimeline?.parameters;
assert.equal(parameters?.length, 2);
assert.equal(parameters?.[0]?.toString(), "config");
assert.equal(parameters?.[1]?.toString(), "cycles");
```

`generateTimeline(config1, 1)` 應該傳回 `[5, 7, 11]`。

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` 應該傳回 `[5, 7, 11, 16, 18, 22]`。

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` 應該傳回 `[3, 1, 7, 10, 8, 14]`。

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` 應該傳回 `[5, 7, 13]`。

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` 應該傳回 `[]`。

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
