---
id: 587d78aa367417b2b2512aed
title: HTML 문서의 Doctype 선언하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cra98AJ'
forumTopicId: 301095
dashedName: declare-the-doctype-of-an-html-document
---

# --description--

지금까지의 도전 과제는 특정 HTML 요소와 그 용도에 대해 다뤘습니다. 하지만 페이지의 전체 구조를 결정하는 몇 가지 요소가 있으며, 모든 HTML 문서에 반드시 포함되어야 합니다.

문서 맨 위에서 브라우저에 페이지가 사용하는 HTML 버전을 알려야 합니다. HTML은 계속 발전하는 언어로 정기적으로 업데이트됩니다. 대부분 주요 브라우저는 최신 명세인 HTML5를 지원하지만, 이전 버전을 사용하는 오래된 웹 페이지도 있습니다.

이 정보를 브라우저에 알리려면 첫 번째 줄에 `<!DOCTYPE ...>` 태그를 추가하는데, 여기서 `...` 부분이 HTML 버전을 나타냅니다. HTML5의 경우 `<!DOCTYPE html>`를 사용합니다.

`!`와 대문자 `DOCTYPE`는 특히 오래된 브라우저에서 중요합니다. `html`는 대소문자를 구분하지 않습니다.

다음으로, 나머지 HTML 코드는 `html` 태그로 감싸야 합니다. 여는 `<html>`는 `<!DOCTYPE html>` 바로 아래에 위치하고, 닫는 `</html>`는 페이지 끝에 위치합니다.

다음은 페이지 구조 예시입니다. 두 `html` 태그 사이에 HTML 코드를 작성하세요.

```html
<!DOCTYPE html>
<html>

</html>
```

# --instructions--

코드 편집기의 빈 HTML 문서 맨 위에 HTML5용 `DOCTYPE` 태그를 추가하세요. 그 아래에는 `html` 여는 태그와 닫는 태그를 추가하고, 그 안에 `h1` 요소를 감싸세요. 제목에는 어떤 텍스트도 포함할 수 있습니다.

# --hints--

코드에 `<!DOCTYPE html>` 태그가 포함되어야 합니다.

```js
assert.match(code,/<!DOCTYPE\s+?html\s*?>/gi);
```

`html` 요소가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('html'), 1);
```

`html` 태그는 하나의 `h1` 요소를 감싸야 합니다.

```js
assert.match(code,/<html>\s*?<h1>\s*?.*?\s*?<\/h1>\s*?<\/html>/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <h1> Hello world </h1>
</html>
```
