---
id: 5c6c06847491271903d37cfd
title: 라디오 버튼과 체크박스에 value 속성 사용하기
challengeType: 0
forumTopicId: 301099
dashedName: use-the-value-attribute-with-radio-buttons-and-checkboxes
---

# --description--

폼이 제출되면 데이터가 서버로 전송되고 선택한 옵션에 대한 항목이 포함됩니다. `radio`와 `checkbox` 유형의 입력은 `value` 속성에서 값을 보고합니다.

예를 들어:

```html
<label for="indoor">
  <input id="indoor" value="indoor" type="radio" name="indoor-outdoor">Indoor
</label>
<label for="outdoor">
  <input id="outdoor" value="outdoor" type="radio" name="indoor-outdoor">Outdoor
</label>
```

여기에는 두 개의 `radio` 입력이 있습니다. 사용자가 `indoor` 옵션을 선택하고 폼을 제출하면 폼 데이터에 `indoor-outdoor=indoor` 라인이 포함됩니다. 이는 "indoor" 입력의 `name`와 `value` 속성에서 가져온 것입니다.

`value` 속성을 생략하면 제출된 폼 데이터는 기본값인 `on`를 사용합니다. 이 경우 사용자가 "indoor" 옵션을 클릭하고 폼을 제출하면 결과 폼 데이터는 `indoor-outdoor=on`가 되는데, 이는 유용하지 않습니다. 따라서 옵션을 식별하기 위해 `value` 속성을 반드시 설정해야 합니다.

# --instructions--

기존의 `radio`와 `checkbox` 입력 각각에 `value` 속성을 추가하세요. 새로운 라디오 버튼이나 체크박스 요소를 만들지 마세요. 입력 레이블 텍스트를 소문자로 변환한 값을 속성 값으로 사용하세요.

# --hints--

라디오 버튼 중 하나는 `value` 속성에 `indoor` 값을 가져야 합니다.

```js
const indoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "indoor");
assert.notEmpty(indoorRadioButton)
```

라디오 버튼 중 하나는 `value` 속성에 `outdoor` 값을 가져야 합니다.

```js
const outdoorRadioButton = [...document.querySelectorAll('label > input[type="radio"]')].filter(x => x.value === "outdoor");
assert.notEmpty(outdoorRadioButton);
```

체크박스 중 하나는 `value` 속성에 `loving` 값을 가져야 합니다.

```js
const lovingCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "loving");
assert.notEmpty(lovingCheckbox); 
```

체크박스 중 하나는 `value` 속성에 `lazy` 값을 가져야 합니다.

```js
const lazyCheckbox = [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "lazy");
assert.notEmpty(lazyCheckbox); 
```

체크박스 중 하나는 `value` 속성에 `energetic` 값을 가져야 합니다.

```js
const energeticCheckbox =  [...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.value === "energetic");
assert.notEmpty(energeticCheckbox);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label><br>
    <label for="loving"><input id="loving" type="checkbox" name="personality" value="loving"> Loving</label>
    <label for="lazy"><input id="lazy" type="checkbox" name="personality" value="lazy"> Lazy</label>
    <label for="energetic"><input id="energetic" type="checkbox" name="personality" value="energetic"> Energetic</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
