---
id: bad87fee1348bd9aedf08835
title: 체크박스 세트 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cqrkJsp'
forumTopicId: 16821
dashedName: create-a-set-of-checkboxes
---

# --description--

폼에서는 여러 답변이 가능한 질문에 <dfn>체크박스</dfn>를 자주 사용합니다.

체크박스는 `input` 유형입니다.

각 체크박스는 자체 `label` 요소 안에 중첩될 수 있습니다. `input` 요소 안에 `label` 요소를 감싸면 체크박스 입력과 이를 둘러싼 레이블 요소가 자동으로 연결됩니다.

관련된 모든 체크박스 입력은 동일한 `name` 속성을 가져야 합니다.

체크박스 `input`와 해당 `label` 사이의 관계를 명확히 정의하려면 `for` 요소에 `label` 속성을 설정하여 연결된 `id` 요소의 `input` 속성과 일치시키는 것이 모범 사례로 여겨집니다.

체크박스 예시는 다음과 같습니다:

```html
<label for="loving"><input id="loving" type="checkbox" name="personality"> Loving</label>
```

# --instructions--

폼에 세 개의 체크박스 세트를 추가하세요. 각 체크박스는 자체 `label` 요소 안에 중첩되어야 합니다. 세 개 모두 `name`의 `personality` 속성을 공유해야 합니다.

# --hints--

페이지에 세 개의 체크박스 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('input[type="checkbox"]'),3);
```

세 개의 체크박스 각각은 자체 `label` 요소 안에 중첩되어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="checkbox"]:only-child'),3);
```

각 `label` 요소에 닫는 태그가 있는지 확인하세요.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length)
```

체크박스에는 `name`의 `personality` 속성이 있어야 합니다.

```js
assert.lengthOf([...document.querySelectorAll('label > input[type="checkbox"]')].filter(x => x.name === "personality"),3);
```

각 체크박스는 `form` 태그 안에 추가되어야 합니다.

```js
assert.strictEqual(document.querySelector('label').parentNode.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <label for="playful"><input id="playful" type="checkbox" name="personality">Playful</label>
    <label for="lazy"><input id="lazy" type="checkbox" 
name="personality">Lazy</label>
    <label for="evil"><input id="evil" type="checkbox" 
name="personality">Evil</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
