---
id: bad87fee1348bd9aedf08834
title: 라디오 버튼 세트 만들기
challengeType: 0
forumTopicId: 16822
dashedName: create-a-set-of-radio-buttons
---

# --description--

사용자가 여러 선택지 중에서 하나의 답변만 선택하도록 할 질문에 <dfn>라디오 버튼</dfn>을 사용할 수 있습니다.

라디오 버튼은 `input`의 한 종류입니다.

각 라디오 버튼은 자체 `label` 요소 안에 중첩할 수 있습니다. `input` 요소 안에 `label` 요소를 감싸면 라디오 버튼 입력과 이를 둘러싼 레이블 요소가 자동으로 연결됩니다.

관련된 모든 라디오 버튼은 동일한 `name` 속성을 가져야 라디오 버튼 그룹이 만들어집니다. 라디오 그룹을 만들면, 하나의 라디오 버튼을 선택할 때 같은 그룹 내 다른 버튼들은 자동으로 선택 해제되어 사용자가 하나의 답변만 제공하도록 보장합니다.

라디오 버튼 예시는 다음과 같습니다:

```html
<label> 
  <input type="radio" name="indoor-outdoor">Indoor 
</label>
```

`for` 요소에 `label` 속성을 설정하는 것이 모범 사례로 간주되며, 그 값은 `id` 요소의 `input` 속성 값과 일치해야 합니다. 이렇게 하면 보조 기술이 레이블과 관련된 `input` 요소 사이에 연결 관계를 만들 수 있습니다. 예를 들어:

```html
<input id="indoor" type="radio" name="indoor-outdoor">
<label for="indoor">Indoor</label>
```

`input` 요소를 `label` 태그 안에 중첩할 수도 있습니다:

```html
<label for="indoor"> 
  <input id="indoor" type="radio" name="indoor-outdoor">Indoor 
</label>
```

# --instructions--

각각 자체 `label` 요소 안에 중첩된 라디오 버튼 두 개를 폼에 추가하세요. 하나는 `indoor` 옵션을, 다른 하나는 `outdoor` 옵션을 가져야 합니다. 두 버튼 모두 `name` 값을 가진 `indoor-outdoor` 속성을 공유하여 라디오 그룹을 만드세요.

# --hints--

페이지에는 두 개의 `radio` 버튼 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('input[type="radio"]'),2);
```

라디오 버튼에는 `name` 값을 가진 `indoor-outdoor` 속성이 있어야 합니다.

```js
assert.lengthOf([...document.querySelectorAll('input[type="radio"]')].filter(input => input.name === "indoor-outdoor"),2);
```

두 라디오 버튼 요소 각각은 자체 `label` 요소 안에 중첩되어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('label > input[type="radio"]:only-child'),2);
```

각 `label` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/label>/g);
assert.match(code,/<label/g);
assert.strictEqual(code.match(/<\/label>/g).length,code.match(/<label/g).length);
```

라디오 버튼 중 하나에는 레이블 `indoor`가 있어야 합니다.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label =>  label.match(/indoor/gi));
assert.isTrue(hasMatch); 
```

라디오 버튼 중 하나에는 레이블 `outdoor`가 있어야 합니다.

```js
const labelTexts = [...document.querySelectorAll('label')].map(label => label.textContent);
const hasMatch = labelTexts.some(label => label.match(/outdoor/gi));
assert.isTrue(hasMatch); 
```

각 라디오 버튼 요소는 `form` 태그 안에 추가되어야 합니다.

```js
const labelParent = document.querySelector('label').parentNode;
assert.strictEqual(labelParent.tagName,'FORM');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
   <label for="indoor"><input id="indoor" type="radio" name="indoor-outdoor"> Indoor</label>
    <label for="outdoor"><input id="outdoor" type="radio" name="indoor-outdoor"> Outdoor</label><br>
    <input type="text" placeholder="cat photo URL" required>
    <button type="submit">Submit</button>
  </form>
</main>
```
