---
id: bad87fee1348bd9aedf08820
title: 이미지를 링크로 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cRdBnUr'
forumTopicId: 18327
dashedName: turn-an-image-into-a-link
---

# --description--

`a` 요소 안에 요소를 중첩하면 링크로 만들 수 있습니다.

이미지를 `a` 요소 안에 중첩하세요. 예시는 다음과 같습니다:

```html
<a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
```

이미지를 죽은 링크로 만들려면 `#` 요소의 `a` 속성으로 `href`를 사용해야 한다는 점을 기억하세요.

# --instructions--

기존 이미지 요소를 `a` (*앵커*) 요소 안에 넣으세요.

이렇게 하면 이미지 위에 커서를 올렸을 때, 커서 모양이 일반 포인터에서 링크 클릭 포인터로 바뀝니다. 이제 사진이 링크가 되었습니다.

# --hints--

기존 `img` 요소는 `a` 요소 안에 중첩되어야 합니다.

```js
const anchor = document.querySelectorAll('a')[1];
const children = anchor.querySelectorAll("img");
assert.notEmpty(children);
```

`a` 요소는 `href` 값으로 설정된 `#` 속성을 가진 죽은 링크여야 합니다.

```js
const anchor = document.querySelectorAll('a')[1];
const parentHREF = anchor.querySelector("img").parentNode.getAttribute('href');
assert.match(parentHREF,new RegExp('#'));
```

각 `a` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length, code.match(/<a/g).length)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
