---
id: bad87fee1348bd9aede08817
title: 단락 안에 앵커 요소 중첩하기
challengeType: 0
forumTopicId: 18244
dashedName: nest-an-anchor-element-within-a-paragraph
---

# --description--

링크를 다른 텍스트 요소 안에 중첩할 수 있습니다.

```html
<p>
  Here's a <a target="_blank" href="https://www.freecodecamp.org"> link to www.freecodecamp.org</a> for you to follow.
</p>
```

예제를 자세히 살펴보겠습니다. 일반 텍스트는 `p` 요소로 감싸져 있습니다:

```html
<p> Here's a ... for you to follow. </p>
```

다음은 닫는 태그 `<a>`가 필요한 *앵커* 요소 `</a>`입니다:  

```html
<a> ... </a>
```

`target`는 링크를 어디서 열지 지정하는 앵커 태그 속성입니다. 값 `_blank`는 새 탭에서 링크를 열도록 지정합니다. `href`는 링크의 URL 주소를 포함하는 앵커 태그 속성입니다:  

```html
<a href="https://www.freecodecamp.org" target="_blank"> ... </a>
```

`link to www.freecodecamp.org` 요소 안의 텍스트 `a`를 <dfn>앵커 텍스트</dfn>라고 하며, 클릭할 링크를 표시합니다:

```html
<a href=" ... " target="...">link to freecodecamp.org</a>
```

예제의 최종 출력은 다음과 같습니다:  

따라가 볼 수 있는 <a href="https://www.freecodecamp.org" target="_blank">www.freecodecamp.org 링크</a>입니다.

# --instructions--

기존 `a` 요소를 새 `p` 요소 안에 중첩하세요. 새 앵커 태그를 만들지 마세요. 새 단락에는 `View more cat photos`라는 텍스트가 있어야 하며, `cat photos`는 링크이고 나머지는 일반 텍스트입니다.

# --hints--

`a` 요소는 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('a'), 1 );
```

`a` 요소는 "`https://www.freecatphotoapp.com`"에 링크해야 합니다.

```js
assert.lengthOf(document.querySelectorAll('a[href="https://www.freecatphotoapp.com"]'),1);
```

`a` 요소는 `cat photos`라는 앵커 텍스트를 가져야 합니다.

```js
assert.match(document.querySelector('a').textContent,/cat\sphotos/gi)
```

새 `p` 요소를 만들어야 합니다. HTML 코드에는 총 3개 이상의 `p` 태그가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p'),3)
```

`a` 요소는 새 `p` 요소 안에 중첩되어야 합니다.

```js
const anchorParent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode;
assert.strictEqual(anchorParent.tagName,"P")
```

`p` 요소는 `View more ` 텍스트(뒤에 공백 포함)를 가져야 합니다.

```js
const textContent = document.querySelector('a[href="https://www.freecatphotoapp.com"]').parentNode.textContent;
assert.match(textContent,/View\smore\s/gi);
```

`a` 요소는 `View more` 텍스트를 <em>가지면 안 됩니다</em>.

```js
assert.notMatch(document.querySelector('a').textContent,/View\smore/gi); 
```

각 `p` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/p>/g);
assert.match(code,/<p/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

각 `a` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/a>/g);
assert.match(code,/<a/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>View more <a target="_blank" href="https://www.freecatphotoapp.com">cat photos</a></p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
