---
id: bad87fee1348bd9aedf08802
title: HTML 주석 해제
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBmG9T7'
forumTopicId: 18329
dashedName: uncomment-html
---

# --description--

주석은 최종 사용자에게 표시되는 결과 출력에 영향을 주지 않으면서 코드 내에서 다른 개발자에게 의견을 남길 수 있는 방법입니다.

주석은 코드를 완전히 삭제하지 않고 비활성화하는 편리한 방법이기도 합니다.

HTML에서 주석은 `<!--`로 시작해서 `-->`로 끝납니다.

# --instructions--

`h1`, `h2`, `p` 요소의 주석을 해제하세요.

# --hints--

`h1` 요소는 주석을 해제하면 페이지에 보여야 합니다.

```js
assert.notEmpty(document.querySelectorAll('h1'));
```

`h2` 요소는 주석을 해제하면 페이지에 보여야 합니다.

```js
assert.notEmpty(document.querySelectorAll('h2'));
```

`p` 요소는 주석을 해제하면 페이지에 보여야 합니다.

```js
assert.notEmpty(document.querySelectorAll('p'));
```

페이지에 주석 종료 태그(`-->`)가 보이지 않아야 합니다.

```js
const elements = document.querySelectorAll('*');
const potentialComments = Array.from(elements).filter(el => el.textContent.includes('-->'));
assert.notExists(potentialComments[1])
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```
