---
id: bad87fee1348bd9aecf08801
title: HTML5 요소 소개
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cBkZGpt7'
forumTopicId: 301097
dashedName: introduction-to-html5-elements
---

# --description--

HTML5는 더 설명적인 HTML 태그를 도입합니다. 여기에는 `main`, `header`, `footer`, `nav`, `video`, `article`, `section` 등이 포함됩니다.

이 태그들은 HTML에 설명적인 구조를 부여하고, HTML을 더 읽기 쉽게 만들며, 검색 엔진 최적화(SEO)와 접근성에 도움을 줍니다. `main` HTML5 태그는 검색 엔진과 다른 개발자가 페이지의 주요 콘텐츠를 찾는 데 도움을 줍니다.

예시 사용법으로, 두 개의 자식 요소가 중첩된 `main` 요소가 있습니다:

```html
<main> 
  <h1>Hello World</h1>
  <p>Hello Paragraph</p>
</main>
```

**참고:** 많은 새로운 HTML5 태그와 그 이점은 적용된 접근성 섹션에서 다룹니다.

# --instructions--

두 번째 `p` 요소를 만들고 다음의 kitty ipsum 텍스트를 넣으세요: `Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.`

그다음 `main` 요소를 만들고 `p` 요소 안에 두 개의 `main` 요소만 중첩하세요.

# --hints--

Kitty Ipsum 텍스트가 들어간 `p` 요소가 2개 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p'),2);
```

각 `p` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

`p` 요소는 제공된 추가 `kitty ipsum` 텍스트의 처음 몇 단어를 포함해야 합니다.

```js
assert.match(document.querySelectorAll('p')[1].textContent,/Purr\s+jump\s+eat/gi);
```

코드에는 `main` 요소가 하나 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`main` 요소는 두 개의 문단 요소를 자식으로 가져야 합니다.

```js
const main = document.querySelector('main');
const children = main.querySelectorAll("p"); 
assert.lengthOf(children,2);
```

첫 번째 문단 태그 앞에 `main` 여는 태그가 와야 합니다.

```js
assert.match(code,(/<main>\s*?<p>/g));
```

두 번째 닫는 문단 태그 뒤에 `main` 닫는 태그가 와야 합니다.

```js
assert.match(code,(/<\/p>\s*?<\/main>/g));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
