---
id: bad87fee1348bd9aedf08801
title: 단락 요소로 정보 전달하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/ceZ7DtN'
forumTopicId: 18202
dashedName: inform-with-the-paragraph-element
---

# --description--

`p` 요소는 웹사이트에서 단락 텍스트에 가장 적합한 요소입니다. `p`은 "paragraph"의 줄임말입니다.

단락 요소는 다음과 같이 만들 수 있습니다:

```html
<p>I'm a p tag!</p>
```

# --instructions--

`p` 요소 아래에 `h2` 요소를 만들고, 텍스트로 `Hello Paragraph`를 넣으세요.

**참고:** 관례상 모든 HTML 태그는 소문자로 작성합니다. 예를 들어 `<p></p>`가 아니라 `<P></P>`로 작성하세요.

# --hints--

코드에 유효한 `p` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('p'),1);
```

`p` 요소에 텍스트 `Hello Paragraph`가 있어야 합니다.

```js
assert.match(document.querySelector('p').textContent,/hello(\s)+paragraph/gi);
```

`p` 요소에는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/p>/g);
assert.strictEqual(code.match(/<\/p>/g).length,code.match(/<p/g).length);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
```

# --solutions--

```html
<h1>Hello World</h1>
<h2>CatPhotoApp</h2>
<p>Hello Paragraph</p>
```
