---
id: bad87fee1348bd9aedc08830
title: HTML5로 필수 입력 필드 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMd4EcQ'
forumTopicId: 18360
dashedName: use-html5-to-require-a-field
---

# --description--

특정 폼 필드를 필수로 지정할 수 있어서 사용자가 해당 필드를 작성하지 않으면 폼을 제출하지 못하게 할 수 있습니다.

예를 들어, 텍스트 입력 필드를 필수로 만들고 싶다면 `required` 요소 안에 `input` 속성을 추가하기만 하면 됩니다. 이렇게요: `<input type="text" required>`

# --instructions--

텍스트 `input`를 `required` 필드로 만들어서 사용자가 이 필드를 작성하지 않으면 폼을 제출하지 못하게 하세요.

그다음 텍스트를 입력하지 않고 폼을 제출해 보세요. HTML5 폼이 이 필드가 필수임을 어떻게 알려주는지 확인하세요.

# --hints--

텍스트 `input` 요소에 `required` 속성이 있어야 합니다.

```js
assert.isTrue(document.querySelector('input').required);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <form action="https://www.freecatphotoapp.com/submit-cat-photo">
    <input type="text" required placeholder="cat photo URL">
    <button type="submit">Submit</button>
  </form>
</main>
```
