---
id: bad88fee1348bd9aedf08816
title: 앵커 요소로 페이지 내부 섹션에 링크하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cyrDRUL'
forumTopicId: 301098
dashedName: link-to-internal-sections-of-a-page-with-anchor-elements
---

# --description--

`a` (*앵커*) 요소는 웹페이지 내의 다른 섹션으로 이동하는 내부 링크를 만드는 데도 사용할 수 있습니다.

내부 링크를 만들려면 링크의 `href` 속성을 해시 기호 `#`와 내부 링크하려는 요소의 `id` 속성 값으로 할당하세요. 보통 페이지 아래쪽에 위치합니다. 그런 다음 링크 대상 요소에 동일한 `id` 속성을 추가해야 합니다. `id`는 요소를 고유하게 설명하는 속성입니다.

아래는 내부 앵커 링크와 그 대상 요소의 예입니다:

```html
<a href="#contacts-header">Contacts</a>
...
<h2 id="contacts-header">Contacts</h2>
```

사용자가 `Contacts` 링크를 클릭하면 **Contacts** 제목 요소가 있는 웹페이지 섹션으로 이동합니다.

# --instructions--

외부 링크를 내부 링크로 바꾸려면 `href` 속성을 `#footer`로 변경하고 텍스트를 `cat photos`에서 `Jump to Bottom`으로 바꾸세요.

링크된 문서가 새 창 탭에서 열리는 것을 방지하려면 앵커 태그에서 `target="_blank"` 속성을 제거하세요.

그런 다음 페이지 하단의 `id` 요소에 값이 `footer`인 `<footer>` 속성을 추가하세요.

# --hints--

페이지에는 앵커 태그가 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('a'), 1);
```

페이지에는 `footer` 태그가 하나만 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('footer'),1);
```

`a` 태그에는 `href` 속성이 "#footer"로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'), '#footer');
```

`a` 태그에는 `target` 속성이 없어야 합니다.

```js
assert.notExists(document.querySelector('a').getAttribute('target'));
```

`a` 텍스트는 "Jump to Bottom"이어야 합니다.

```js
assert.match(document.querySelector('a').textContent,/Jump to Bottom/gi);
```

`footer` 태그에는 `id` 속성이 "footer"로 설정되어야 합니다.

```js
assert.strictEqual(document.querySelector('footer').getAttribute('id'),'footer');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer>Copyright Cat Photo App</footer>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>

  <a href="#footer">Jump to Bottom</a>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched. Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff. Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
  <p>Meowwww loved it, hated it, loved it, hated it yet spill litter box, scratch at owner, destroy all furniture, especially couch or lay on arms while you're using the keyboard. Missing until dinner time toy mouse squeak roll over. With tail in the air lounge in doorway. Man running from cops stops to pet cats, goes to jail.</p>
  <p>Intently stare at the same spot poop in the plant pot but kitten is playing with dead mouse. Get video posted to internet for chasing red dot leave fur on owners clothes meow to be let out and mesmerizing birds leave fur on owners clothes or favor packaging over toy so purr for no reason. Meow to be let out play time intently sniff hand run outside as soon as door open yet destroy couch.</p>

</main>

<footer id="footer">Copyright Cat Photo App</footer>
```
