---
id: bad87fee1348bd9aedf08812
title: 웹사이트에 이미지 추가하기
challengeType: 0
forumTopicId: 16640
dashedName: add-images-to-your-website
---

# --description--

웹사이트에 이미지를 추가하려면 `img` 요소를 사용하고 `src` 속성으로 특정 이미지의 URL을 지정하세요.

예시는 다음과 같습니다:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg">
```

`img`가 빈 요소라는 점을 참고하세요.

모든 `img` 요소에는 반드시 `alt` 속성이 있어야 합니다. `alt` 속성 안의 텍스트는 화면 읽기 프로그램의 접근성을 향상시키기 위해 사용되며, 이미지가 로드되지 않을 때 표시됩니다.

**참고:** 이미지가 순수한 장식용이라면 빈 `alt` 속성을 사용하는 것이 모범 사례입니다.

이상적으로 `alt` 속성에는 필요하지 않은 한 특수 문자가 포함되지 않아야 합니다.

위 `alt` 예시에 `img` 속성을 추가해 봅시다:

```html
<img src="https://www.freecatphotoapp.com/your-image.jpg" alt="freeCodeCamp logo">
```

# --instructions--

웹사이트에 이미지를 추가해 봅시다:

기존 `main` 요소 안에, 기존 `img` 요소들 앞에 `p` 요소를 삽입하세요.

그다음 `src` 속성을 설정하여 URL `https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg` 를 가리키게 하세요.

마지막으로 `img` 요소에 적절한 텍스트가 들어간 `alt` 속성을 꼭 추가하세요.

# --hints--

페이지에 이미지 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('img'));
```

이미지에는 고양이 이미지 URL을 가리키는 `src` 속성이 있어야 합니다.

```js
const url = document.querySelector('img').getAttribute('src');
assert.match(url,/^https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg$/i);
```

이미지 요소의 `alt` 속성은 비어 있으면 안 됩니다.

```js
assert.exists(document.querySelector('img').getAttribute('alt'));
assert.isNotEmpty(document.querySelector('img').getAttribute('alt'));
assert.match(__helpers.removeWhiteSpace(code),/<(?:img|IMG)\S*alt=(['"])(?!\1|>)\S+\1\S*\/?>/)
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>


  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
