---
id: bad87fee1348bd9aedf08817
title: 해시 기호로 작동하지 않는 링크 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cMdkytL'
forumTopicId: 18230
dashedName: make-dead-links-using-the-hash-symbol
---

# --description--

때때로 링크가 어디로 연결될지 모를 때 `a` 요소를 웹사이트에 미리 추가하고 싶을 때가 있습니다.

나중에 배울 `JavaScript`을 사용해 링크 동작을 변경할 때도 이 방법이 유용합니다.

# --instructions--

현재 `href` 속성 값은 "`https://www.freecatphotoapp.com`"를 가리키는 링크입니다. 작동하지 않는 링크를 만들려면 `href` 속성 값을 해시 기호라고도 하는 `#`로 바꾸세요.

예를 들어: `href="#"`

# --hints--

`a` 요소는 `href` 속성 값이 "#"인 작동하지 않는 링크여야 합니다.

```js
assert.strictEqual(document.querySelector('a').getAttribute('href'),'#');
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="https://www.freecatphotoapp.com" target="_blank">cat photos</a>.</p>

  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#" target="_blank">cat photos</a>.</p>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
