---
id: bad87fee1348bd9aedf08804
title: HTML 주석 처리하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cGyGbca'
forumTopicId: 16782
dashedName: comment-out-html
---

# --description--

주석을 시작하려면 `<!--`를 사용하고, 주석을 끝내려면 `-->`를 사용해야 한다는 점을 기억하세요.

여기서는 `h2` 요소가 시작되기 전에 주석을 끝내야 합니다.

# --instructions--

`h1` 요소와 `p` 요소는 주석 처리하고, `h2` 요소는 주석 처리하지 마세요.

# --hints--

`h1` 요소는 페이지에 보이지 않도록 주석 처리해야 합니다.

```js
assert.isEmpty(document.querySelectorAll('h1'));
```

`h2` 요소는 페이지에 보이도록 주석 처리하지 않아야 합니다.

```js
assert.isNotEmpty(document.querySelectorAll('h2'));
```

`p` 요소는 페이지에 보이지 않도록 주석 처리해야 합니다.

```js
assert.isEmpty(document.querySelectorAll('p'));
```

각 주석은 `-->`로 닫아야 합니다.

```js
assert.isAbove(code.match(/[^fc]-->/g).length, 1);
```

코드 내에서 `h1`, `h2`, `p` 요소의 순서는 변경하지 마세요.

```js
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[0],'<h1>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[1],'<h2>');
assert.strictEqual(code.match(/<([a-z0-9]){1,2}>/g)[2],'<p>');
```

# --seed--

## --seed-contents--

```html
<!--
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
-->
```

# --solutions--

```html
<!--<h1>Hello World</h1>-->
<h2>CatPhotoApp</h2> 
<!--<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p> -->
```
