---
id: bad87fee1348bd9aedf08830
title: 텍스트 필드에 자리 표시자 텍스트 추가하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cKdJDhg'
forumTopicId: 16647
dashedName: add-placeholder-text-to-a-text-field
---

# --description--

자리 표시자 텍스트는 사용자가 아무 입력도 하기 전 `input` 요소에 표시되는 내용입니다.

자리 표시자 텍스트는 다음과 같이 만들 수 있습니다:

```html
<input type="text" placeholder="this is placeholder text">
```

**참고:** `input`이 빈 요소라는 점을 기억하세요.

# --instructions--

텍스트 `placeholder`의 `input` 값을 "cat photo URL"로 설정하세요.

# --hints--

기존 텍스트 `placeholder` 요소에 `input` 속성을 추가해야 합니다.

```js
assert.notEmpty(document.querySelectorAll('input[placeholder]'));
```

`placeholder` 속성의 값을 `cat photo URL`로 설정해야 합니다.

```js
assert.exists(document.querySelector('input'));
assert.exists(document.querySelector('input').getAttribute('placeholder'));
const placeholder = document.querySelector('input').getAttribute('placeholder');
assert.match(placeholder,/cat\s+photo\s+URL/gi);
```

완성된 `input` 요소에는 닫는 태그가 없어야 합니다.

```js
assert.notMatch(code,/<input.*\/?>.*<\/input>/gi);
```

완성된 `input` 요소는 올바른 구문을 가져야 합니다.

```js
assert.notEmpty(document.querySelectorAll('input[type=text]'));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text">
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>
  
  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
  
  <p>Things cats love:</p>
  <ul>
    <li>catnip</li>
    <li>laser pointers</li>
    <li>lasagna</li>
  </ul>
  <p>Top 3 things cats hate:</p>
  <ol>
    <li>flea treatment</li>
    <li>thunder</li>
    <li>other cats</li>
  </ol>
  <input type="text" placeholder="cat photo URL">
</main>
```
