---
id: bad87fee1348bd9aedf08833
title: 자리 표시자 텍스트로 빈칸 채우기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cgR7Dc7'
forumTopicId: 18178
dashedName: fill-in-the-blank-with-placeholder-text
---

# --description--

웹 개발자는 전통적으로 자리 표시자 텍스트로 <dfn>lorem ipsum 텍스트</dfn>를 사용합니다. lorem ipsum 텍스트는 고대 로마의 키케로가 쓴 유명한 구절에서 무작위로 발췌한 것입니다.

lorem ipsum 텍스트는 16세기부터 활자 조판사들이 자리 표시자 텍스트로 사용해 왔으며, 이 전통은 웹에서도 계속되고 있습니다.

하지만 5세기는 충분히 긴 시간입니다. 우리가 CatPhotoApp을 만들고 있으니, "kitty ipsum" 텍스트를 사용해 보겠습니다.

# --instructions--

`p` 요소 안의 텍스트를 이 kitty ipsum 텍스트의 처음 몇 단어로 바꾸세요: `Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.`

# --hints--

`p` 요소에는 제공된 "kitty ipsum" 텍스트의 처음 몇 단어가 들어 있어야 합니다.

```js
assert.match(document.querySelector('p').textContent,/Kitty(\s)+ipsum/gi);
```

# --seed--

## --seed-contents--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Hello Paragraph</p>
```

# --solutions--

```html
<h1>Hello World</h1>

<h2>CatPhotoApp</h2>

<p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff</p>
```
