---
id: bad87fee1348bd9aedf08827
title: 순서 없는 목록 만들기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/cDKVPuv'
forumTopicId: 16814
dashedName: create-a-bulleted-unordered-list
---

# --description--

HTML에는 <dfn>순서 없는 목록</dfn> 또는 글머리 기호 스타일 목록을 만드는 특별한 요소가 있습니다.

순서 없는 목록은 `<ul>` 요소로 시작하고, 그 뒤에 여러 개의 `<li>` 요소가 옵니다. 마지막으로, 순서 없는 목록은 `</ul>`로 닫힙니다.

예를 들어:

```html
<ul>
  <li>milk</li>
  <li>cheese</li>
</ul>
```

는 `milk`와 `cheese`의 글머리 기호 스타일 목록을 만듭니다.

# --instructions--

마지막 두 개의 `p` 요소를 제거하고 페이지 하단에 고양이가 좋아하는 세 가지를 순서 없는 목록으로 만드세요.

# --hints--

`ul` 요소를 만드세요.

```js
assert.isNotEmpty(document.querySelectorAll('ul'));
```

`li` 요소 안에 세 개의 `ul` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('ul li'),3);
```

`ul` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/ul>/gi);
assert.match(code,/<ul/gi);
assert.strictEqual(code.match(/<\/ul>/gi).length,code.match(/<ul/gi).length);
```

`li` 요소들도 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/li>/gi);
assert.match(code,/<li[\s>]/gi);
assert.strictEqual(code.match(/<\/li>/gi).length,code.match(/<li[\s>]/gi).length);
```

`li` 요소는 빈 문자열이나 공백만 포함해서는 안 됩니다.

```js
assert.isEmpty([...document.querySelectorAll('ul li')].filter((item) => item.textContent.trim() === ""));
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  <p>Click here to view more <a href="#">cat photos</a>.</p>

  <a href="#"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

  <ul>
    <li>milk</li>
    <li>mice</li>
    <li>catnip</li>
  </ul>
</main>
```
