---
id: bad87fee1348bd9aedf08816
title: 앵커 요소로 외부 페이지 링크하기
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVMPUv/c8EkncB'
forumTopicId: 18226
dashedName: link-to-external-pages-with-anchor-elements
---

# --description--

웹 페이지 외부의 콘텐츠로 링크하려면 `a` (*앵커*) 요소를 사용할 수 있습니다.

`a` 요소는 `href` 속성이라고 하는 목적지 웹 주소가 필요합니다. 또한 앵커 텍스트도 필요합니다. 예시는 다음과 같습니다:

```html
<a href="https://www.freecodecamp.org">this links to freecodecamp.org</a>
```

그러면 브라우저는 `this links to freecodecamp.org` 텍스트를 클릭할 수 있는 링크로 표시합니다. 그리고 그 링크는 `https://www.freecodecamp.org` 웹 주소로 이동합니다.

# --instructions--

`a` 요소를 만들어 `https://www.freecatphotoapp.com` 로 링크하고 앵커 텍스트는 "cat photos"로 하세요.

# --hints--

`a` 요소의 앵커 텍스트는 `cat photos` 여야 합니다.

```js
assert.match(document.querySelector('a').textContent,/cat photos/gi);
```

`a` 요소가 `https://www.freecatphotoapp.com` 로 링크해야 합니다.

```js
assert.match(document.querySelector('a').getAttribute('href'),/^https?:\/\/(www\.)?freecatphotoapp\.com\/?$/i);
```

`a` 요소는 닫는 태그가 있어야 합니다.

```js
assert.match(code,/<\/a>/g);
assert.strictEqual(code.match(/<\/a>/g).length,code.match(/<a/g).length);
```

# --seed--

## --seed-contents--

```html
<h2>CatPhotoApp</h2>
<main>



  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">

  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```

# --solutions--

```html
<h2>CatPhotoApp</h2>
<main>
  
  <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
  
  <a href="https://www.freecatphotoapp.com">cat photos</a>
  <p>Kitty ipsum dolor sit amet, shed everywhere shed everywhere stretching attack your ankles chase the red dot, hairball run catnip eat the grass sniff.</p>
  <p>Purr jump eat the grass rip the couch scratched sunbathe, shed everywhere rip the couch sleep in the sink fluffy fur catnip scratched.</p>
</main>
```
