---
id: 61537bb9b1a29430ac15ad38
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inicie este proyecto vinculando su archivo `styles.css` a la página.

# --hints--

Su elemento de `link` debe tener un atributo `href` con el valor `styles.css`.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
