---
id: 68c18b1d0fce081f85312d97
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Katika funzo lililopita, ulijifunza kwamba sifa ya kipengele yenye kiambishi awali kimoja cha mstari wa chini hutumiwa kwa matumizi ya ndani kwa mkataba.

Badilisha sifa za `name` na `level` kuwa `_name` na `_level`, kwa kuwa hizi hazipaswi kubadilishwa kutoka nje ya darasa lao.

Kumbuka kwamba hili halizuizi sifa ya kipengele kufikiwa au kubadilishwa kutoka nje ya darasa. Pia, katika Python daima kuna njia ya kufikia sifa za kipengele binafsi (zilizo na kiambishi awali cha mistari miwili ya chini).

# --hints--

Njia yako ya `__init__` inapaswa kuweka `self._name` kuwa `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._name = name")`)) })
```

Njia yako ya `__init__` inapaswa kuweka `self._level` kuwa `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self._level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

charlie_brown = Employee('Charlie Brown', 'trainee')
--fcc-editable-region--
```
