---
id: 68c16924498728d259e4781c
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Sasa tengeneza mfano wa darasa la `Employee` kwa kupitisha mfuatano wa herufi `Charlie Brown` na `trainee`. Weka thamani ya mfano huo katika kigezo kinachoitwa `charlie_brown`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("charlie_brown")`)) })
```

Unapaswa kuweka thamani ya `Employee('Charlie Brown', 'trainee')` kwa `charlie_brown`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("charlie_brown").is_equivalent("charlie_brown = Employee('Charlie Brown', 'trainee')")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self.name = name
        self.level = level

--fcc-editable-region--

--fcc-editable-region--

```
