---
id: 68c15a2a1f0e70ca7154aba5
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaanzisha mfumo wa kufuatilia mshahara wa wafanyakazi.

Anza kwa kuunda darasa la vitu linaloitwa `Employee`. Ndani yake unda njia ya `__init__` yenye vigezo `self`, `name`, na `level`. Ndani ya njia ya `__init__`, weka thamani `name` na `level` kwa sifa za kipengele za mfano zenye majina sawa.

# --hints--

Unapaswa kuwa na darasa la vitu linaloitwa `Employee`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Employee")`)) })
```

Darasa lako la `Employee` linapaswa kuwa na njia ya `__init__` yenye vigezo vitatu `self`, `name`, na `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_args("self, name, level")`)) })
```

Njia yako ya `__init__` inapaswa kuweka `self.name` kuwa `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Njia yako ya `__init__` inapaswa kuweka `self.level` kuwa `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.level = level")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

```
