---
id: 68c28d659d345270d4a61084
title: Hatua 11
challengeType: 20
dashedName: step-11
---

# --description--

Njia ya `__repr__` ni njia maalum inayotarajiwa kurudisha mfuatano wa herufi unaowakilisha kitu ambacho kinaweza kutumika kuanzisha kitu hicho.

Kwa mfano, njia ya `__repr__` ya `Employee('Charlie', 'developer')` inapaswa kurudisha mfuatano wa herufi `Employee('Charlie', 'developer')`, ambao ni mfuatano huo huo unaotumika kuunda kitu hicho.

Toa darasa la `Employee` lenye njia ya `__repr__` yenye kigezo `self`, na lifanye lirudishe mfuatano wa herufi unaoweza kutumika kuanzisha kitu hicho.

# --hints--

Darasa lako la `Employee` linapaswa kuwa na njia ya `__repr__` yenye kigezo `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__repr__").has_args("self")`)) })
```

Njia yako ya `__repr__` inapaswa kurudisha mfuatano wa herufi unaotumika kuanzisha kitu hicho.

```js
({ test: () => runPython(`
  empl = Employee("Frank", "dreamer")
  assert repr(empl) == 'Employee("Frank", "dreamer")' or repr(empl) == "Employee('Frank', 'dreamer')"
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
