---
id: 68c28c6517e7de5e85f8ebe5
title: Hatua 9
challengeType: 20
dashedName: step-9
---

# --description--

Sasa kwa kuwa una wapataji kwa sifa za kipengele `_name` na `_level`, sasisha mfuatano wa herufi unaorudishwa na `__str__` kutumia `self.name` na `self.level`. Hii itafanya wito kwa wapataji badala ya kupata sifa za kipengele moja kwa moja.

# --hints--

Njia yako ya `__str__` inapaswa kurudisha `f'{self.name}: {self.level}'`

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_return("f'{self.name}: {self.level}'")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self._name}: {self._level}'
--fcc-editable-region--
    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(charlie_brown.name)
print(charlie_brown.level)
```
