---
id: 68c27eb292bfec6b9402d58f
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Ongeza njia ya `__str__` katika darasa la vitu la `Employee`. Ifanye irudishe mfuatano wa herufi wa f-string wenye muundo wa `name: level`, ukibadilisha `name` na `level` kwa sifa za kipengele zinazofanana.

Baada ya hapo, chapisha `charlie_brown` kwenye konsoli.

# --hints--

Darasa lako la vitu la `Employee` linapaswa kuwa na njia ya `__str__` yenye kigezo cha `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__str__").has_args("self")`)) })
```

Njia yako ya `__str__` inapaswa kurudisha mfuatano wa herufi wa f-string wenye muundo wa `name: level`, ukibadilisha `name` na `level` kwa thamani za `self._name` na `self._level`, mtawalia.

```js
({ test: () => runPython(`
    import io
    import sys
        
    captured_output = io.StringIO()
    sys.stdout = captured_output
        
    empl = Employee("Frank", "dreamer")
    print(empl)
        
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
        
    assert "Frank: dreamer" in output
`) })
```

Unapaswa kuchapisha `charlie_brown` kwenye konsoli.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

--fcc-editable-region--
    

charlie_brown = Employee('Charlie Brown', 'trainee')

--fcc-editable-region--

```
