---
id: 68c281cecda94ba2c6838e17
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Sasa umefafanua getter kwa ajili ya `name`, unaweza kupata sifa ya kipengele `_name` kupitia sifa `name`. Kwa hiyo chapisha `charlie_brown.name` kwenye konsoli.

# --hints--

Unapaswa kuchapisha `charlie_brown.name` kwenye konsoli.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(charlie_brown.name)")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    def __init__(self, name, level):
        self._name = name
        self._level = level

    def __str__(self):
        return f'{self._name}: {self._level}'

    @property
    def name(self):
        return self._name

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
--fcc-editable-region--

--fcc-editable-region--

```
