---
id: 68c7f7012a700243eff1cbc0
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Kiwango kinachotumika kuanzisha mfanyakazi kinapaswa kuchaguliwa kati ya viwango maalum. Utatumia `_base_salaries` kuthibitisha kiwango na kuweka mshahara sahihi kwa mfanyakazi.

Baada ya tamko lako la `if` lililopo, tengeneza tamko lingine la `if` linalokagua kama `level` haipo katika `Employee._base_salaries`.

Ndani ya tamko jipya la `if`, ondoa `ValueError` yenye ujumbe wa `Invalid value '{level}' for 'level' attribute.`, ambapo `{level}` inapaswa kubadilishwa na thamani ya hoja ya `level`.

# --hints--

Unapogundua kuwa `level` haipo katika `Employee._base_salaries`, unapaswa kutoa `ValueError` yenye ujumbe wa `Invalid value '{level}' for 'level' attribute.`, ambapo `{level}` inapaswa kubadilishwa na thamani ya hoja ya `level`.

```js
({ test: () => runPython(`
  try:
    Employee("Frank", "dreamer")
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
    assert False, "Expected to raise ValueError with invalid level"
`) })
```

Haufai kutoa kosa lolote wakati `level` iko katika `Employee._base_salaries`.

```js
({ test: () => runPython(`
  levels = [
    "trainee",
    "junior",
    "mid-level",
    "senior"
  ]
  for level in levels:
    try:
      Employee("Frank", level)
    except Exception:
      assert False, f"Expected not to raise ValueError with level '{level}'"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
--fcc-editable-region--
    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        
        self._name = name
        self._level = level
--fcc-editable-region--
    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @property
    def level(self):
        return self._level

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
```
