---
id: 68c8002beb157fe47d2e7699
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Setter ni njia inayotumika kuweka thamani ya sifa, ikiruhusu ukaguzi wa uthibitisho na vikwazo. Unaweza kuunda setter kwa kutumia dekoreta ya `@propertyName.setter`, ambapo `propertyName` inapaswa kuendana na jina la sifa ya kuweka:

```py
class Person:
    def __init__(self, name):
        self.name = name  # Calling the setter

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, value):
        self._name = value

p = Person('Alice')
p.name = 'Abigail' # Calls the setter
print(p.name) # Abigail
```

Baada ya njia yako ya getter, tengeneza njia ya `name` yenye vigezo `self` na `new_name`. Pamba njia hiyo kwa `@name.setter`. Ndani ya njia, weka `self._name` kuwa `new_name`.

# --hints--

Darasa lako la `Employee` linapaswa kuwa na njia ya `name` yenye vigezo `self` na `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_args("self, new_name")`)) })
```

Njia yako ya `name` inapaswa kupambwa na `@name.setter`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_decorators("name.setter")`)) })
```

Njia yako ya `name` inapaswa kuweka `self._name` kuwa `new_name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].has_stmt("self._name = new_name")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

--fcc-editable-region--
    
--fcc-editable-region--

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
