---
id: 68c809c990f253912a9e9209
title: Hatua 22
challengeType: 20
dashedName: step-22
---

# --description--

Kama ulivyojifunza katika funzo lililopita, setter hutoa njia ya kudhibiti jinsi sifa ya kipengele inaweza kubadilishwa. Ili kuhakikisha kuwa `new_name` ni aina sahihi, tengeneza tamko la `if` linaloinua `TypeError` lenye ujumbe `'name' must be a string.` wakati `new_name` si mfano wa `str`.

# --hints--

Unapaswa kuwa na tamko la `if` ndani ya setter yako ya `name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("name")[1].find_ifs()[0]`)) })
```

Unapaswa kuinua `TypeError` lenye ujumbe `'name' must be a string.` wakati `new_name` si mfano wa `str`.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  non_names = [0, True, []]
  for i in non_names:
    try:
      emp.name = i
    except TypeError as e:
      assert str(e) == "'name' must be a string."
    else:
      assert False, "Expected to raise TypeError with non-string new_name"
`) })
```

Hupaswi kuinua kosa lolote wakati `new_name` ni mfuatano wa herufi.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.name = "Jack"
  except Exception:
    assert False, "Expected not to raise any exception with valid new_name"
    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
--fcc-editable-region--
        
--fcc-editable-region--
        self._name = new_name

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
