---
id: 68c9c3aa714bc326e026b826
title: Hatua 30
challengeType: 20
dashedName: step-30
---

# --description--

Kiwango kipya hakiwezi kuwekwa bila kuangalia kama ni kiwango halali. Baada ya ukaguzi wa `isinstance`, tengeneza usemi wa `if` unaoinua `ValueError` wakati `new_level` si ufunguo wa `Employee._base_salaries`.

Kwa ujumbe wa kosa, tumia `Invalid value '{new_level}' for 'level' attribute.`, ambapo `{new_level}` inapaswa kubadilishwa na hoja iliyotumwa kwa setter.

# --hints--

Unapaswa kuwa na usemi wa `if` ndani ya setter yako ya `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[1]`)) })
```

Wakati `new_level` si ufunguo wa `Employee._base_salaries`, unapaswa kuinua `ValueError` na ujumbe `Invalid value '{new_level}' for 'level' attribute.`, ambapo `{new_level}` inapaswa kubadilishwa na hoja iliyotumwa kwa setter.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "dreamer"
  except ValueError as e:
    assert str(e) == "Invalid value 'dreamer' for 'level' attribute."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Haupaswi kuinua kosa lolote wakati `new_level` ni ufunguo wa `Employee._base_salaries`

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  for i in Employee._base_salaries:
    try:
      emp.level = i
    except Exception:
      assert False, "Expected not to raise any exception with valid new_level"    
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
        
    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
