---
id: 68c9a09163d4a42e9bc9b638
title: Hatua 24
challengeType: 20
dashedName: step-24
---

# --description--

Sasa una njia ya kuweka (setter method) kwa sifa ya kipengele `_name` ambayo inafungua sifa `name`.

Jaribu sasa kusahihisha `charlie_brown` `name` kwa jina tofauti kama unavyofanya kawaida kwa sifa yoyote. Hii itasababisha wito wa njia ya kuweka (setter method) kufanyika kwa siri.

# --hints--

Unapaswa kusahihisha `charlie_brown` `name` kwa jina tofauti.

```js
({ test: () => runPython(`
  assert _Node(_code).has_variable("charlie_brown.name")
  assert charlie_brown.name != 'Charlie Brown'
`) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        if not isinstance(level, str):
            raise TypeError("'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self.name = name
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
--fcc-editable-region--

--fcc-editable-region--

```
