---
id: 68c9cab4b1118da59eecfc56
title: Hatua 31
challengeType: 20
dashedName: step-31
---

# --description--

Baada ya tamko la `if` lililopo, tengeneza lingine ili kuinua `ValueError` wakati `new_level` tayari ni kiwango kilichochaguliwa.

Kumbuka kuwa `_level` bado haipo wakati wa uanzishaji, kwa hivyo tumia `hasattr(self, '_level')` kuangalia kama ipo kabla ya kulinganisha. Hii inazuia `AttributeError` wakati kitu kinapotengenezwa kwa mara ya kwanza.

Kwa ujumbe, tumia `'{level}' is already the selected level.`, ambapo `{level}` inapaswa kubadilishwa na kiwango cha sasa.

# --hints--

Unapaswa kuwa na tamko la tatu la `if` ndani ya setter yako ya `level`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2]`)) })
```

Wakati `new_level` ni sawa na `self.level`, unapaswa kuinua `ValueError` na ujumbe wa `'{level}' is already the selected level.`, ambapo `{level}` inapaswa kubadilishwa na kiwango cha sasa.

```js
({ test: () => runPython(`
  emp = Employee('Frank', 'trainee')
  try:
    emp.level = "trainee"
  except ValueError as e:
    assert str(e) == "'trainee' is already the selected level."
  else:
      assert False, "Expected to raise ValueError with invalid new_level"
`) })
```

Tamko lako la `if` linapaswa kutumia `hasattr(self, '_level')` kuangalia kama `_level` ipo kabla ya kulinganisha.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_functions("level")[1].find_ifs()[2].find_conditions()[0].is_equivalent("hasattr(self, '_level') and new_level == self.level")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }

    def __init__(self, name, level):
        self.name = name
        self.level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        if not isinstance(new_name, str):
            raise TypeError("'name' must be a string.")
        self._name = new_name
        print(f"'name' updated to '{self.name}'.")

    @property
    def level(self):
        return self._level

    @level.setter
    def level(self, new_level):
        if not isinstance(new_level, str):
            raise TypeError("'level' must be a string.")
        if new_level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{new_level}' for 'level' attribute.")
--fcc-editable-region--
        
--fcc-editable-region--
        self._level = new_level
    

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')
```
