---
id: 69b014adbc624861e6ff3dd1
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Sasa baada ya kuundwa kwa setter ya `name`, sahihisha `__init__` ili itumie `self.name = name` badala ya `self._name = name`. Hii inawezesha mgawo wa thamani wa `name` kupitia setter, na uthibitisho utaongezwa katika hatua inayofuata.

Pia, rekebisha sharti la `if` ili likague tu `level`, na sasisha ujumbe wa `TypeError` ipasavyo:

# --hints--

Njia yako ya `__init__` inapaswa kutumia `self.name = name` badala ya `self._name = name`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").has_stmt("self.name = name")`)) })
```

Njia yako ya `__init__` inapaswa kugagua tu `isinstance(level, str)` na kusasisha ujumbe wa `TypeError` ipasavyo.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Employee").find_function("__init__").find_if("not isinstance(level, str)").find_body().has_stmt("raise TypeError(\\"'level' attribute must be of type 'str'.\\")")`)) })
```

# --seed--

## --seed-contents--

```py
class Employee:
    _base_salaries = {
        'trainee': 1000,
        'junior': 2000,
        'mid-level': 3000,
        'senior': 4000,
    }
    
    def __init__(self, name, level):
--fcc-editable-region--
        if not (isinstance(name, str) and isinstance(level, str)):
            raise TypeError("'name' and 'level' attribute must be of type 'str'.")
        if level not in Employee._base_salaries:
            raise ValueError(f"Invalid value '{level}' for 'level' attribute.")
        self._name = name
--fcc-editable-region--
        self._level = level
        self._salary = Employee._base_salaries[level]

    def __str__(self):
        return f'{self.name}: {self.level}'

    def __repr__(self):
        return f"Employee('{self.name}', '{self.level}')"

    @property
    def name(self):
        return self._name

    @name.setter
    def name(self, new_name):
        self._name = new_name
        
    @property
    def level(self):
        return self._level

    @property
    def salary(self):
        return self._salary

charlie_brown = Employee('Charlie Brown', 'trainee')
print(charlie_brown)
print(f'Base salary: ${charlie_brown.salary}')

```
