---
id: 673b567e3ba535dda140d278
title: Ein Wahlsystem erstellen
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

In diesem Labor erstellen Sie ein Wahlsystem, das `Map` verwendet, um eine Umfrage zu erstellen, und `Set`, um doppelte Stimmen zu verhindern.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `poll`-Variable mit einem neuen `Map`-Objekt initialisieren.

2. Sie sollten eine Funktion `addOption` haben, die einen Parameter `option` akzeptiert.

3. In der Funktion `addOption`:

   - Wenn `option` noch nicht in der Umfrage existiert, sollte sie mit einem leeren `Set` als Wert hinzugefügt werden, um die Wähler zu verfolgen. Sie sollten außerdem die Nachricht `Option "<option>" added to the poll.` zurückgeben.

   - Wenn `option` bereits existiert, sollte die Nachricht `Option "<option>" already exists.` zurückgegeben werden.

   - Wenn Sie versuchen, eine leere Option hinzuzufügen, sollte die Funktion die Nachricht `Option cannot be empty.` zurückgeben.

4. Sie sollten eine Funktion `vote` haben, die zwei Parameter akzeptiert: `option` (die Option, für die abgestimmt wird) und `voterId` (eine eindeutige ID des Wählers).

5. In der Funktion `vote`:

   - Wenn `option` nicht in der Umfrage existiert, sollte die Funktion die Nachricht `Option "<option>" does not exist.` zurückgeben.

   - Wenn `option` existiert, sollte die Funktion prüfen, ob `voterId` bereits für diese `option` abgestimmt hat.

   - Wenn der Wähler bereits abgestimmt hat, sollte die Funktion die Nachricht `Voter <voterId> has already voted for "<option>".` zurückgeben.

   - Wenn der Wähler noch nicht abgestimmt hat, sollte `voterId` zum `Set` von `voters` für diese Option hinzugefügt werden. Die Funktion sollte die Nachricht `Voter <voterId> voted for "<option>".` zurückgeben.

6. Sie sollten mindestens drei Optionen in Ihrem `poll` haben.

7. Ihr `poll` sollte mindestens drei Stimmen haben.

8. Sie sollten eine Funktion `displayResults` haben, die die Umfrageergebnisse im folgenden Format zurückgibt:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

Sie sollten eine `poll`-Variable haben, die mit einem neuen `Map`-Objekt initialisiert ist.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

Sie sollten eine Funktion `addOption` definieren, die einen Parameter `option` akzeptiert.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

Sie sollten eine Funktion `vote` definieren, die zwei Parameter `option` und `voterId` akzeptiert.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

Sie sollten eine Funktion `displayResults` definieren, um die Umfrageergebnisse anzuzeigen.

```js
assert.isFunction(displayResults);
```

Sie sollten mindestens drei Optionen in Ihrem `poll` haben.

```js
assert.isAtLeast(poll.size, 3);
```

Ihr `poll` sollte mindestens drei Stimmen haben.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

Sie sollten sicherstellen, dass jede Wahloption auf ein `Set`-Objekt abgebildet wird.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` sollte `Option "Egypt" added to the poll.` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Das Hinzufügen einer leeren Option sollte `"Option cannot be empty."` zurückgeben.

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

Wenn `Turkey` bereits hinzugefügt wurde, sollte `addOption("Turkey")` `Option "Turkey" already exists.` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wenn `Malaysia` in den Wahloptionen existiert, sollte `vote("Malaysia", "traveler1")` `Voter traveler1 voted for "Malaysia".` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` sollte die `Set` der Wähler für eine Option aktualisieren.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wenn `traveler1` versucht, erneut für `Algeria` abzustimmen, sollte `vote("Algeria", "traveler1")` `Voter traveler1 has already voted for "Algeria".` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Doppelte Stimmen sollten die Größe von `Set` nicht erhöhen.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wenn `Nigeria` nicht in den Wahloptionen ist, sollte `vote("Nigeria", "traveler2")` `Option "Nigeria" does not exist.` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Eine eindeutige Option sollte mehrere Stimmen erhalten können.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

Wenn Option `Turkey` zwei Stimmen erhält, Option `Morocco` eine Stimme und Option `Spain` keine Stimme, sollte `displayResults()` `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"` zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` sollte die Ergebnisse im korrekten Format zurückgeben.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
