---
id: 65576ff7888f9e96f52a4be1
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Los diccionarios almacenan datos en forma de pares *clave*-*valor*. Una clave está separada del valor correspondiente por dos puntos. Y cada par clave-valor está separado del siguiente par por una coma:

```py
my_dict = {
    'name': 'Michael',
    'occupation': 'Lumberjack'
}
```

Añade un nuevo par clave-valor al diccionario. Use the string `'species'` as the key, and the string `'guinea pig'` as the value.

# --hints--

You should have a `'species'` key with the value `'guinea pig'` inside your `copper` dictionary.

```js
({ test: () => assert(runPython(`
    copper == {"species": "guinea pig"}
  `))
})
```

Your `copper` dictionary should have a single key-value pair.

```js
({ test: () => assert(runPython(`
    len(copper) == 1
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
copper = {}
--fcc-editable-region--
```
