---
id: 65578d4fc3afc3b8f554c882
title: Paso 41
challengeType: 20
dashedName: step-41
---

# --description--

`min()` puede recibir también una palabra clave como argumento. Al pasar una función como argumento adicional a `min()` puedes modificar el método que compara los elementos de la lista.

El resultado de la línea que acabas de escribir en el paso anterior es el primer nodo por orden alfabético. En su lugar quieres seleccionar el nodo no visitado situado a menor distancia desde el nodo inicial.

Pasa `key=distances.get` como segundo argumento en tu llamada a `min()`. De esta manera la comparación se llevará a cabo dependiendo del valor que cada elemento de la lista `unvisited` tenga dentro del diccionario `distances`.

# --hints--

Deberías pasar `key=distances.get` como segundo argumento en tu llamada a `min()`.

```js
({ test: () =>  {
    const shortest = __helpers.python.getDef(code, "shortest_path");
    const {function_body} = shortest;    
    assert(function_body.match(/^(\s{4})while\s+unvisited\s*:\s*^\1\1current\s*=\s*min\s*\(\s*unvisited\s*,\s*key\s*=\s*distances\.get\s*\)/ms));
  }
})
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--    
    while unvisited:
        current = min(unvisited)
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
