---
id: 65578d0f6c78a0b868a43b9c
title: Paso 40
challengeType: 20
dashedName: step-40
---

# --description--

Dentro del bucle `while`, lo primero que hay que hacer es definir el nodo actual a visitar. Para eso puedes usar la función `min()`. Devuelve el elemento más pequeño del iterable pasado como argumento.

Elimina la palabra reservada `pass`, luego crea una variable llamada `current` y asígnala a `min(unvisited)`.

# --hints--

Debes crear una variable `current` en el bucle `while`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].has_variable("current")`)) })
```

You should assign `min(unvisited)` to your `current` variable. Remember to delete `pass`.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("shortest_path").find_whiles()[0].find_bodies()[0].is_equivalent("current = min(unvisited)")`)) })
```

# --seed--

## --seed-contents--

```py
my_graph = {
    'A': [('B', 3), ('D', 1)],
    'B': [('A', 3), ('C', 4)],
    'C': [('B', 4), ('D', 7)],
    'D': [('A', 1), ('C', 7)]
}

def shortest_path(graph, start):
    unvisited = list(graph)
    distances = {node: 0 if node == start else float('inf') for node in graph}
    paths = {node: [] for node in graph}
    paths[start].append(start)
--fcc-editable-region--
    while unvisited:
        pass
--fcc-editable-region--    
    print(f'Unvisited: {unvisited}\nDistances: {distances}\nPaths: {paths}')

#shortest_path(my_graph, 'A')

```
