---
id: 687d7deea6a2a22fa8816a73
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У попередніх уроках ви дізналися про деякі загальні структури даних. У цьому практичному занятті ви реалізуєте зв’язаний список.

Спочатку створіть клас із назвою `LinkedList`, щоб представити структуру, яка триматиме вузли вашого списку. Додайте ключове слово `pass` всередині класу, щоб тест пройшов успішно.

# --hints--

Ви маєте створити клас із назвою `LinkedList`. Не забудьте додати ключове слово `pass`.

```js
({
    test: () => runPython(`assert _Node(_code).has_class("LinkedList")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
