---
id: 687d8e05a529a4838f1c961c
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Далі видаліть ключове слово `pass` всередині класу `Node`, а потім додайте класу `Node` метод `__init__` з двома параметрами, `self` та `element`.

Всередині методу призначте `element` значенню `self.element`. Це зберігає значення, передане під час створення вузла.

# --hints--

Ви повинні визначити метод `__init__` всередині свого класу `Node`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").has_function("__init__")`) 
})
```

Ваш метод `__init__` має мати два параметри, `self` та `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_class("Node").find_function("__init__").has_args("self, element")`) 
})
```

Ви повинні призначити `element` значенню `self.element` всередині методу `__init__`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.element == "A"
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
--fcc-editable-region--
    class Node:
        pass
--fcc-editable-region--
```
