---
id: 687d8c6f066877768f7ee441
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Всередині вашого класу `LinkedList` замініть ключове слово `pass` на внутрішній клас під назвою `Node`. Він буде представляти окремі вузли, які зберігають дані у списку.

# --hints--

Ви маєте створити клас `Node` всередині класу `LinkedList`. Не забудьте додати ключове слово `pass` всередині класу `Node`.

```js
({
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").has_class("Node")`) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class LinkedList:
    pass
--fcc-editable-region--
```
