---
id: 687e344cc1993f7316af030c
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Всередині класу `LinkedList` додайте метод `__init__`, який встановлює `self.length` у значення `0`. Це ініціалізує фішку для відстеження кількості вузлів (елементів), які наразі є у списку.

# --hints--

Ви повинні встановити `self.length` у `0` всередині методу `__init__` класу `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.length == 0`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    
--fcc-editable-region--
```
