---
id: 687e366b3004a9831e4cbd41
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Додайте метод з назвою `is_empty`, який повертає `self.length == 0`.

Ви будете використовувати це, щоб відстежувати, скільки вузлів у списку. Отже, з `self.length == 0` він повертає `True`, що означає, що список не має вузлів.

Якщо `self.length` більше за `0`, він повертає `False`, що означає, що список має принаймні один вузол.

# --hints--

Ви маєте створити метод `is_empty`, який повертає `self.length == 0`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.is_empty() == True
    ll.length = 1
    assert ll.is_empty() == False`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

--fcc-editable-region--
    
--fcc-editable-region--
```
