---
id: 68800a64b06e80d9a31c131b
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ще в методі `__init__` класу `LinkedList` встановіть `self.head` у `None`. Це слугуватиме референсом на перший вузол у списку. Наразі список порожній, тому потрібно встановити його в `None`.


# --hints--

Ви повинні встановити `self.head` у `None` всередині методу `__init__` класу `LinkedList`.

```js
({ 
    test: () => runPython(`
    ll = LinkedList()
    assert ll.head is None`) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None

--fcc-editable-region--
    def __init__(self):
        self.length = 0
--fcc-editable-region--
```
