---
id: 688003f589b875af60273dd9
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ще всередині методу `__init__` призначте `None` значенню `self.next`. Це встановить вказівник `next` на `None`, що означає, що новий вузол поки не вказує на жоден інший вузол у списку.

Пізніше `next` можна оновити, щоб він вказував на інший вузол і формував ланцюг зв’язаного списку.

# --hints--

Всередині методу `__init__` слід ініціалізувати атрибут `next` значенням `None`.

```js
({ 
    test: () => runPython(`
    n = LinkedList.Node("A") 
    assert n.next is None
    `) 
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
--fcc-editable-region--
        def __init__(self, element):
            self.element = element
--fcc-editable-region--
```
