---
id: 68857e0a49baf4751f737a2a
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Тепер, після методу `is_empty`, створіть метод під назвою `add` всередині класу `LinkedList`. Новий метод має мати два параметри, `self` і `element`.

Всередині вашого нового методу створіть змінну з ім’ям `node` і призначте їй новий екземпляр `Node`, використовуючи `self.Node(element)`. Це створить новий вузол, який буде додано до зв’язаного списку.

# --hints--

Ви повинні створити метод `add` всередині класу `LinkedList`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").has_function("add")`
      )
    )
})
```

Ваш метод `add` має мати два параметри, `self` і `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("add").has_args("self, element")`) 
})
```

Ви повинні призначити `self.Node(element)` `node` у методі `add`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_body().is_equivalent("node = self.Node(element)")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
