---
id: 68867c3898b09115d8c5140f
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Цикл `while`, який ви щойно створили, дає змогу пройти через кожен вузол у списку, доки не буде досягнуто останнього вузла. Коли це станеться, ви можете додати вузол до списку.

У блоці `else`, одразу після блоку `while`, призначте `node` `current_node.next`.

# --hints--

Ви повинні призначити `node` `current_node.next` у блоці `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].has_stmt("current_node.next = node")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
