---
id: 688649c8673f842f49228a5e
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Якщо список не порожній, метод має знайти останній вузол і зв’язати новий вузол з його вказівником `next`.

Для цього створіть блок `else` після блоку `if`. Усередині блоку `else` створіть змінну `current_node` і призначте їй `self.head`.

# --hints--

Ви маєте призначити `self.head` `current_node` усередині блоку `else`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("add").find_ifs()[0].find_bodies()[1].is_equivalent("current_node = self.head")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
--fcc-editable-region--
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())
```
