---
id: 6888d05ba074461b64f8eaca
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Далі ви створите ще один метод, який буде використовуватися для видалення заданого елемента зі зв’язаного списку.

Створіть метод з назвою `remove` з двома параметрами, `self` та `element`. Усередині нього створіть змінну з назвою `previous_node` і призначте їй `None`.

# --hints--

Ви повинні створити метод з назвою `remove`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_class("LinkedList").has_function("remove")
    `)
  )
})
```

Ваш метод `remove` має мати два параметри, `self` та `element`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("LinkedList").find_function("remove").has_args("self, element")`) 
})
```

Ви повинні призначити `None` до `previous_node` у методі `remove`.

```js
({
  test: () =>
    assert(
      runPython(
        `_Node(_code).find_class("LinkedList").find_function("remove").find_body().is_equivalent("previous_node = None")`
      )
    )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
