---
id: 6888fe6c745b6852ed240229
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Далі ви маєте пройтися циклом по зв’язаному списку, щоб знайти вузол, який містить вказаний елемент для видалення.

Створіть `while` цикл, який перевіряє, коли `current_node` не є `None` і коли `current_node.element` не дорівнює `element`. Всередині `while` циклу призначте `current_node` значенню `previous_node`.

# --hints--

У вас має бути `while` цикл всередині `remove` методу.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
`)})
```

Ваш `while` цикл має перевіряти, коли `current_node` не є `None` і коли `current_node.element` не дорівнює `element`.

```js
({
  test: () => runPython(`
  _cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0].find_conditions()[0]
  assert _cond.is_equivalent("current_node is not None and current_node.element != element")
`)})
```

Ви маєте призначити `current_node` значенню `previous_node` всередині `while` циклу.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_whiles()[0]
assert _cond.is_equivalent("while current_node is not None and current_node.element != element:\\n previous_node = current_node")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

--fcc-editable-region--
    def remove(self, element):
        previous_node = None
        current_node = self.head
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
