---
id: 688acb2053ac7978a06345f1
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Далі створіть інструкцію `elif`, яка перевіряє, що `previous_node` не дорівнює `None`. Усередині неї призначте `current_node.next` значенню `previous_node.next`.

Завдяки цьому, якщо елемент для видалення знайдено і він не знаходиться у головному вузлі, вказівник `next` попереднього вузла буде оновлено так, щоб пропустити поточний вузол (елемент для видалення), видаляючи його зі зв’язаного списку.

# --hints--

Ви повинні створити інструкцію `elif`, яка перевіряє, що `previous_node` не дорівнює `None`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0].find_conditions()[1]
assert _cond.is_equivalent("previous_node is not None")
`) })
```

Ви повинні призначити `current_node.next` значенню `previous_node.next` всередині інструкції `elif`.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_class("LinkedList").find_function("remove").find_ifs()[0]
assert _cond.is_equivalent("if current_node is None:\\n return\\nelif previous_node is not None:\\n previous_node.next = current_node.next")
`) })
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

    def remove(self, element):
        previous_node = None
        current_node = self.head
        while current_node is not None and current_node.element != element:
            previous_node = current_node
            current_node = current_node.next
--fcc-editable-region--
        if current_node is None:
            return
--fcc-editable-region--

my_list = LinkedList()
print(my_list.is_empty())

my_list.add(1)
my_list.add(2)
print(my_list.is_empty())
print(my_list.length)
```
