---
id: 68b84df2648b4c96b25022f3
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Давайте перевіримо ваш зв’язаний список на цей момент! Внизу вашого коду створіть екземпляр `LinkedList`, який назвіть `my_list`. Потім, використовуючи метод `is_empty`, який ви створили на попередньому кроці, виведіть `my_list.is_empty()`.

# --hints--

Ви маєте створити екземпляр `LinkedList`, який назвете `my_list`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_variable("my_list").value_is_call("LinkedList")
    `)
  )
})
```

Ви маєте вивести `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).has_call("print(my_list.is_empty())")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0

--fcc-editable-region--

--fcc-editable-region--
```
