---
id: 68b9e951e4892721632acc26
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Чудово! Тепер давайте перевіримо додавання елементів до вашого зв’язаного списку. Внизу вашого коду, після створення екземпляра `LinkedList`, названого `my_list`, і виведення `my_list.is_empty()`, додайте `1` і `2` до списку за допомогою методу `add`.

Потім знову виведіть `my_list.is_empty()`, щоб переконатися, що список більше не порожній. Нарешті, виведіть `my_list.length`, щоб побачити поточну довжину зв’язаного списку.

# --hints--

Ви повинні додати `1` до списку за допомогою методу `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(1)")`)
  )
})
```

Ви повинні додати `2` до списку за допомогою методу `add`.

```js
({
  test: () => assert(
    runPython(`_Node(_code).has_call("my_list.add(2)")`)
  )
})
```

Ви повинні вивести `my_list.is_empty()`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[1].is_equivalent("print(my_list.is_empty())")
    `)
  )
})
```

Ви повинні вивести `my_list.length`.

```js
({
  test: () => assert(
    runPython(`
      _Node(_code).find_calls("print")[2].is_equivalent("print(my_list.length)")
    `)
  )
})
```

# --seed--

## --seed-contents--

```py
class LinkedList:
    class Node:
        def __init__(self, element):
            self.element = element
            self.next = None
            
    def __init__(self):
        self.length = 0
        self.head = None

    def is_empty(self):
        return self.length == 0
    
    def add(self, element):
        node = self.Node(element)
        if self.is_empty():
            self.head = node
        else:
            current_node = self.head
            while current_node.next is not None:
                current_node = current_node.next
            current_node.next = node
        self.length += 1

my_list = LinkedList()
print(my_list.is_empty())

--fcc-editable-region--

--fcc-editable-region--
```
