---
id: 65c4f013851cefd1a4fe4c96
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

The `__init__` method takes two parameters: `self` (which represents the instance of the class being created) and `key` (the value to be stored in the node). Fügen Sie diese beiden Parameter zur `__init__()` Methode hinzu.

# --hints--

Sie sollten den `Self-` Parameter zur Methode hinzufügen.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Sie sollten den `-Taste` Parameter zur Methode hinzufügen.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
