---
id: 61329b210dac0b08047fd6ab
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

Continuing with the `meta` elements, a `viewport` definition tells the browser how to render the page. Including one betters visual accessibility on mobile, and improves _SEO_ (search engine optimization).

添加帶有 `content` 屬性的 `viewport` 定義，詳細說明頁面的 `width` 和 `initial-scale`。

# --hints--

你應該在 `head` 中創建另一個 `meta` 元素。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

你應該給 `meta` 元素添加一個 `name` 屬性，值爲 `viewport`。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

你應該給 `meta` 添加一個 `content` 屬性，值爲 `width=device-width, initial-scale=1`。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
