---
id: 6133acc353338c0bba9cb553
title: 步驟5
challengeType: 0
dashedName: step-5
---

# --description--

Lastly in the `head`, the `title` element is useful for screen readers to understand the content of a page. Furthermore, it is an important part of _SEO_.

給你的頁面一個描述性的簡潔的 `title`。

# --hints--

你應該給 `head` 元素添加一個 `title` 元素。

```js
assert.exists(document.querySelector('head > title'));
```

`title` 的文本不應超過 60 個字符。

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

嘗試讓你的 `title` 元素的文本作更多描述性說明。 _提示：至少 15 個字符_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
