---
id: 69b279bae0f20eda63163cf6
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

In un passaggio precedente, hai assegnato l'elemento nella `leftPart` alla lista `array` perché era più piccolo. Ma questo non sarà sempre il caso. In alcuni casi di confronto, l'elemento a destra potrebbe essere più piccolo.

Crea una clausola `else` da eseguire quando l'elemento in `leftPart` non è minore dell'elemento in `rightPart`.

All'interno del blocco `else`, assegna `rightPart[rightArrayIndex]` a `array[sortedIndex]`.

# --hints--

Dovresti creare un blocco `else` che contenga `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
