---
id: 69b279bae0f20eda63163cf9
title: Passaggio 15
challengeType: 1
dashedName: step-15
---

# --description--

Il loop `while` che hai creato confronta un elemento da `leftPart` con un altro in `rightPart`, poi aggiunge l'elemento più piccolo alla lista principale `array`.

Continuerà questa operazione finché non ci saranno più elementi da confrontare. Ma `leftPart` potrebbe ancora avere elementi mentre `rightPart` non ne ha, e viceversa.

Crea un altro loop `while` che venga eseguito quando `leftArrayIndex` è minore della lunghezza di `leftPart`. Nei passaggi successivi, lo userai per copiare gli elementi rimanenti in `leftPart` nella lista `array`.

# --hints--

Dovresti creare un loop `while` con la condizione `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
