---
id: 69b279bae0f20eda63163cf8
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Le istruzioni `if` e `else` che hai creato nei passaggi precedenti assegneranno gli elementi all'array ordinato.

Ogni elemento assegnato all'array ordinato occupa un indice nella lista. Quindi devi passare all'indice successivo nell'array ordinato dopo ogni assegnazione.

Sotto il blocco `if`/`else`, ma ancora all'interno del loop `while`, incrementa `sortedIndex` di `1`.

# --hints--

Dovresti incrementare il valore corrente di `sortedIndex` di `1` dopo la clausola `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
