---
id: 69b278d82f0e514f32acb724
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Dopo aver tagliato la metà sinistra, ti serve ancora la metà destra per completare il passaggio di unione.

Usa la sintassi slice per estrarre la metà destra di `array` e assegnarla a una variabile chiamata `rightPart`.

# --hints--

Dovresti avere una variabile `rightPart` all'interno della funzione `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Dovresti usare la sintassi slice per estrarre la metà destra dell'array e assegnarla alla variabile `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
