---
id: 69b278519ef46b2fd8bf1546
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Hai già un valore `middlePoint`. Usalo per dividere l'array in due metà e tieni pronta la porzione sinistra per l'ordinamento.

Usa la sintassi slice per estrarre la metà sinistra di `array` e assegnala a una variabile chiamata `leftPart`.

# --hints--

Dovresti dichiarare una variabile `leftPart` all'interno della funzione `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Dovresti assegnare `array.slice(0, middlePoint)` a `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
