---
id: 69b278d82f0e514f32acb724
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Dopo aver tagliato la metà sinistra, ti serve ancora la metà destra per completare il passaggio di unione.

Usa la sintassi slice per estrarre la metà destra di `array` e assegnarla a una variabile chiamata `rightPart`.

# --hints--

Dovresti avere una variabile `rightPart` all'interno della funzione `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
```

Dovresti usare la sintassi slice per estrarre la metà destra dell'array e assegnarla alla variabile `rightPart`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
assert.isTrue(
  rightPart.value.matches('array.slice(middlePoint)'),
  'You should assign array.slice(middlePoint) to rightPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
