---
id: 69b279bae0f20eda63163cf1
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora è il momento di ordinare e unire le liste (`leftPart` e `rightPart`) nella lista originale `array`.

Puoi farlo confrontando gli elementi di entrambe le liste e unendo l'elemento più piccolo nella lista principale. Farai questo confronto per tutti gli indici in `leftPart` e `rightPart`.

Crea tre variabili: `leftArrayIndex`, `rightArrayIndex` e `sortedIndex`, e imposta i loro valori a `0`. Queste variabili ti aiuteranno a tenere traccia di ogni indice durante il processo di ordinamento.

# --hints--

Dovresti dichiarare una variabile `leftArrayIndex` inizializzata a `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

Dovresti dichiarare una variabile `rightArrayIndex` inizializzata a `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

Dovresti dichiarare una variabile `sortedIndex` inizializzata a `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
