---
id: 69b279bae0f20eda63163cf4
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Quando la condizione `if` valuta a `true`, significa che l'elemento nella lista `leftPart` è più piccolo dell'elemento con cui viene confrontato nella lista `rightPart`.

In quel caso, puoi mettere l'elemento trovato in `leftArrayIndex` in `leftPart` all'interno dell'array ordinato.

All'interno del blocco `if`, assegna `leftPart[leftArrayIndex]` a `array[sortedIndex]`.

# --hints--

Dovresti assegnare `leftPart[leftArrayIndex]` a `array[sortedIndex]` dentro `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
