---
id: 69b279bae0f20eda63163cf3
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

All'interno del loop, inizia a confrontare gli elementi correnti di ciascuna metà così puoi unire prima il valore più piccolo.

Aggiungi un blocco `if` che verifichi se l'elemento corrente in `leftPart` è minore dell'elemento corrente in `rightPart`. 
# --hints--

La condizione `if` dovrebbe verificare se `leftPart[leftArrayIndex]` è minore di `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
