---
id: 69b279bae0f20eda63163cf1
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora è il momento di ordinare e unire le liste (`leftPart` e `rightPart`) nella lista originale `array`.

Puoi farlo confrontando gli elementi di entrambe le liste e unendo l'elemento più piccolo nella lista principale. Farai questo confronto per tutti gli indici in `leftPart` e `rightPart`.

Crea tre variabili: `leftArrayIndex`, `rightArrayIndex` e `sortedIndex`, e imposta i loro valori a `0`. Queste variabili ti aiuteranno a tenere traccia di ogni indice durante il processo di ordinamento.

# --hints--

Dovresti dichiarare una variabile `leftArrayIndex` inizializzata a `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Dovresti dichiarare una variabile `rightArrayIndex` inizializzata a `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Dovresti dichiarare una variabile `sortedIndex` inizializzata a `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
