---
id: 69b2713ff6c21fed2d3ad88b
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort è un algoritmo di ordinamento che usa il principio divide et impera per ordinare collezioni di dati. In pratica, "divide" una collezione in sotto-parti più piccole, e "conquista" le sotto-parti ordinandole indipendentemente, poi unisce le sotto-parti ordinate.

Inizia creando una funzione chiamata `mergeSort` che prende il parametro `array`.

# --hints--

Dovresti dichiarare una funzione chiamata `mergeSort` con un singolo parametro `array`.

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
