---
id: 69b2713ff6c21fed2d3ad88b
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

Il merge sort è un algoritmo di ordinamento che usa il principio divide et impera per ordinare collezioni di dati. In pratica, "divide" una collezione in sotto-parti più piccole, e "conquista" le sotto-parti ordinandole indipendentemente, poi unisce le sotto-parti ordinate.

Inizia creando una funzione chiamata `mergeSort` che prende il parametro `array`.

# --hints--

Dovresti dichiarare una funzione chiamata `mergeSort` con un singolo parametro `array`.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
