---
id: 69b279bae0f20eda63163cf0
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Alla fine del corpo della tua funzione, richiama di nuovo la funzione `mergeSort()`. Questa volta, passa `rightPart` come argomento alla chiamata della funzione.

# --hints--

Dovresti chiamare `mergeSort()` con l'argomento `rightPart` alla fine del corpo della tua funzione.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
