---
id: 69b279bae0f20eda63163cf2
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Successivamente, devi creare un loop che continui finché ci sono elementi rimanenti sia in `leftPart` sia in `rightPart`.

Per farlo, crea un loop `while` con due condizioni: una che verifica se `leftArrayIndex` è minore della lunghezza di `leftPart` e un'altra condizione che verifica se `rightArrayIndex` è minore della lunghezza di `rightPart`.

# --hints--

Dovresti creare un loop `while` che controlla se `leftArrayIndex` è minore della lunghezza di `leftPart` e se `rightArrayIndex` è minore della lunghezza di `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
