---
id: 69b2762b17d23fbee40a6c00
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

L'algoritmo di ordinamento merge sort esegue tre azioni:

- Dividere una sequenza di elementi non ordinata in sotto-parti
- Ordinare gli elementi nelle sotto-parti
- Unire le sotto-parti ordinate

Quanto sopra avviene in modo ricorsivo finché le sotto-parti non vengono unite nella sequenza completa ordinata. Inizia dividendo la sequenza.

Per prima cosa, all'interno della funzione `mergeSort`, dichiara una variabile `middlePoint` usando `const` e assegnale `Math.floor(array.length / 2)` così puoi dividere la sequenza a metà prima di ricorrere.

# --hints--

Dovresti dichiarare una variabile `middlePoint` e assegnarle `Math.floor(array.length / 2)`.

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
