---
id: 69b279bae0f20eda63163cfe
title: Passaggio 20
challengeType: 1
dashedName: step-20
---

# --description--

Prima di testare la funzione `mergeSort()`, devi creare un caso base che fermi l'esecuzione della funzione quando la lunghezza di `array` è minore o uguale a `1`.

Questo caso base fermerà la chiamata ricorsiva. Senza di esso, l'operazione di ordinamento per unione continuerebbe a funzionare anche quando la lista è già ordinata o non contiene elementi.

Subito dopo la dichiarazione della funzione, crea un'istruzione `if` che verifichi se `array.length` è minore o uguale a 1. All'interno del corpo dell'istruzione `if`, aggiungi un'istruzione `return` per fermare l'esecuzione della funzione.

# --hints--

Dovresti creare un'istruzione `if` con la condizione `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Dovresti avere un'istruzione `return` nel corpo dell'istruzione `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
