---
id: 69b279bae0f20eda63163cef
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora che hai diviso la lista `array` in due liste separate, continuerai a dividere ogni lista finché ogni elemento non sarà da solo nella propria lista. Una lista con un solo numero è sempre ordinata.

Per farlo, chiama ricorsivamente `mergeSort` all'interno della tua funzione e passa `leftPart` come argomento alla chiamata.

# --hints--

Dovresti chiamare `mergeSort(leftPart)` all'interno della stessa funzione `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
