---
id: 6989f48c3f0292c84f8d8938
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Jetzt müssen Sie erkennen, wann Sie eine vollständige Kombination von runden Klammern erstellt haben. Eine vollständige Kombination hat eine Länge, die dem Doppelten der Anzahl der Paare entspricht (da jedes Paar eine öffnende und eine schließende Klammer beiträgt).

Fügen Sie innerhalb Ihrer `while`-Schleife eine `if`-Anweisung hinzu, um zu prüfen, ob `current.length === 2 * pairs` zutrifft. Wenn diese Bedingung wahr ist, fügen Sie `current` dem `result`-Array hinzu.

# --hints--

Sie sollten eine `if`-Anweisung mit der Bedingung `current.length === 2 * pairs` innerhalb Ihrer `while`-Schleife haben.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Sie sollten `current` dem `result`-Array innerhalb Ihrer neuen `if`-Anweisung hinzufügen.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
