---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Jetzt müssen Sie das Hinzufügen von schließenden runden Klammern behandeln. Die wichtigste Regel ist, dass Sie eine schließende Klammer nur hinzufügen dürfen, wenn dadurch das Gleichgewicht erhalten bleibt, das heißt, es müssen mehr öffnende runde Klammern verwendet worden sein als schließende.

Fügen Sie innerhalb des `else`-Blocks eine zweite `if`-Anweisung hinzu, um zu prüfen, ob `closesUsed` kleiner ist als `opensUsed`.

Wenn diese Bedingung wahr ist, fügen Sie ein weiteres neues Array zur `queue`: `[current + ')', opensUsed, closesUsed + 1]` hinzu. Dies stellt den Zustand nach dem Hinzufügen einer schließenden Klammer dar.

# --hints--

Sie sollten eine zweite `if`-Anweisung mit der Bedingung `closesUsed < opensUsed` innerhalb Ihres `else`-Blocks haben.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

Sie sollten `[current + ')', opensUsed, closesUsed + 1]` zur Warteschlange hinzufügen, wenn die Bedingung wahr ist.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
