---
id: 6989f207306332c06b0add39
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Für den Breitensuche-Ansatz verwenden Sie eine Warteschlange, um verschiedene Zustände zu verfolgen, während Sie die Kombinationen der runden Klammern aufbauen. Jeder Zustand wird als Array dargestellt, das drei Elemente enthält:

- Der aktuell aufgebaute String
- Die Anzahl der bisher verwendeten öffnenden runden Klammern
- Die Anzahl der bisher verwendeten schließenden runden Klammern

Erstellen Sie eine Variable namens `queue` und initialisieren Sie sie mit einem Array, das ein Array enthält: `['', 0, 0]`. Dies stellt den Anfangszustand mit einem leeren String und null verwendeten Klammern dar.

# --hints--

Sie sollten eine Variable namens `queue` in Ihrer `genParentheses`-Funktion deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

Sie sollten `queue` mit einem Array initialisieren, das das Array `['', 0, 0]` enthält.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
