---
id: 6989ef885db3a3488a13b426
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Bevor Sie den Kernalgorithmus implementieren, müssen Sie die Eingabe validieren. Der `pairs`-Parameter sollte eine ganze Zahl sein, da er die Anzahl der zu generierenden Klammerpaare darstellt.

Fügen Sie am Anfang Ihrer Funktion eine `if`-Anweisung hinzu, um zu prüfen, ob `pairs` keine ganze Zahl ist. Verwenden Sie dafür die `Number.isInteger()`-Methode.

Wenn die Bedingung zutrifft, geben Sie den String `The number of pairs should be an integer` zurück.

# --hints--

Sie sollten eine if-Anweisung mit der Bedingung `!Number.isInteger(pairs)` innerhalb der `genParentheses`-Funktion haben.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Sie sollten einen String zurückgeben, wenn die Eingabe keine ganze Zahl ist.

```js
let result = genParentheses("i");
assert.isString(result);
```

Sie sollten den String `The number of pairs should be an integer` zurückgeben, wenn die Eingabe ein String ist.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```

