---
id: 6989f555728e63ad8b349d84
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Wenn der aktuelle String noch nicht vollständig ist, müssen Sie die nächsten möglichen Zustände erkunden. Fügen Sie eine `else`-Klausel für die aktuelle `if`-Anweisung hinzu.

Innerhalb des `else`-Blocks behandeln Sie das Hinzufügen von öffnenden runden Klammern. Fügen Sie eine `if`-Anweisung hinzu, um zu prüfen, ob `opensUsed` kleiner als `pairs` ist. So stellen Sie sicher, dass Sie nicht mehr öffnende runde Klammern verwenden, als erlaubt sind.

Wenn diese Bedingung `true` ist, fügen Sie ein neues Array in die `queue` ein: `[current + '(', opensUsed + 1, closesUsed]`. Dies stellt den Zustand nach dem Hinzufügen einer öffnenden runden Klammer dar.

# --hints--

Sie sollten eine `else`-Klausel für Ihre `if`-Anweisung innerhalb der `while`-Schleife haben.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

Sie sollten eine `if`-Anweisung innerhalb des `else`-Blocks mit der Bedingung `opensUsed < pairs` haben.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

Sie sollten `[current + '(', opensUsed + 1, closesUsed]` in die Warteschlange einfügen, wenn die Bedingung `true` erfüllt ist.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
