---
id: 6989f1ba358acba651d5ee6b
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt richten Sie die Datenstruktur ein, um Ihre Ergebnisse zu speichern. Erstellen Sie eine Variable namens `result` und initialisieren Sie sie mit einem leeren Array. Dieses Array speichert alle gültigen Kombinationen von runden Klammern, die Sie erzeugen.

Aktualisieren Sie Ihre return-Anweisung, sodass `result` anstelle eines leeren Arrays zurückgegeben wird.

# --hints--

Sie sollten eine Variable namens `result` in Ihrer `genParentheses`-Funktion deklarieren.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.result);
```

Sie sollten `result` mit einem leeren Array initialisieren.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { result } = genParentheses?.variables ?? {};
assert.isFalse(result?.value.isEmpty());
const resultValue = eval(result?.value.toString());
assert.isArray(resultValue);
assert.isEmpty(resultValue);
```

Sie sollten am Ende Ihrer Funktion `result` zurückgeben.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.isTrue(genParentheses?.hasReturn('result'));
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
