---
id: 6989f2e35e6941e13e445156
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Jetzt implementieren Sie die Haupt-BFS-Schleife. Erstellen Sie eine `while`-Schleife, die so lange läuft, wie die `queue` nicht leer ist, das heißt, `queue.length` größer als `0` ist.

Geben Sie innerhalb der Schleife `queue` in der Konsole aus.

# --hints--

Sie sollten eine `while`-Schleife mit der Bedingung `queue.length > 0` erstellen.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Sie sollten `queue` innerhalb Ihrer `while`-Schleife in der Konsole ausgeben.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
--fcc-editable-region--
  
--fcc-editable-region--
  return result;
}
```
