---
id: 6989eb65324d97775c4965fe
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop implementieren Sie eine Funktion, die alle gültigen Kombinationen von runden Klammern mithilfe eines Breitensuche-Ansatzes (BFS) erzeugt. Zum Beispiel sind die gültigen Kombinationen von zwei Paaren runder Klammern `(())` und `()()`.

Beginnen Sie damit, eine Funktion namens `genParentheses` mit einem einzelnen Parameter `pairs` zu erstellen. Geben Sie vorerst ein leeres Array aus der Funktion zurück.

# --hints--

Sie sollten eine Funktion namens `genParentheses` definieren.

```js
assert.isFunction(genParentheses);
```

Ihre `genParentheses`-Funktion sollte einen einzelnen Parameter namens `pairs` haben.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Ihre Funktion sollte ein leeres Array zurückgeben.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
