---
id: 6989f10a72ad424d4ad6d514
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Als Nächstes müssen Sie überprüfen, ob die Anzahl der Paare mindestens eins beträgt, da Sie keine Kombinationen von runden Klammern mit null oder negativen Paaren erzeugen können.

Fügen Sie eine weitere `if`-Anweisung hinzu, um zu prüfen, ob `pairs` kleiner als `1` ist. Wenn diese Bedingung zutrifft, geben Sie den String `The number of pairs should be at least 1` zurück.

# --hints--

Sie sollten eine zweite `if`-Anweisung mit der Bedingung `pairs < 1` haben.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Sie sollten den String `The number of pairs should be at least 1` zurückgeben, wenn `pairs` kleiner als `1` ist.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
