---
id: 6989f1ba358acba651d5ee6b
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt richten Sie die Datenstruktur ein, um Ihre Ergebnisse zu speichern. Erstellen Sie eine Variable namens `result` und initialisieren Sie sie mit einem leeren Array. Dieses Array speichert alle gültigen Kombinationen von runden Klammern, die Sie erzeugen.

Aktualisieren Sie Ihre return-Anweisung, sodass `result` anstelle eines leeren Arrays zurückgegeben wird.

# --hints--

Sie sollten eine Variable namens `result` in Ihrer `genParentheses`-Funktion deklarieren.

```js
assert.match(genParentheses.toString(), /result/);
```

Sie sollten `result` mit einem leeren Array initialisieren.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Sie sollten am Ende Ihrer Funktion `result` zurückgeben.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
