---
id: 6989f3632f8b8bdf168aa321
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Innerhalb Ihrer `while`-Schleife verwenden Sie `queue.shift()`, um das erste Element aus der Warteschlange zu entfernen und zu erhalten. Dies implementiert das First-in-first-out-(FIFO-)Verhalten, das für BFS charakteristisch ist.

Zerlegen Sie dieses Array in drei Variablen: `current`, `opensUsed` und `closesUsed`.

# --hints--

Sie sollten `queue.shift()` verwenden, um das erste Element aus der Warteschlange zu erhalten.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Sie sollten das Ergebnis in drei Variablen zerlegen: `current`, `opensUsed` und `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
