---
id: 66c057041df6394ca796bf33
title: 문장 생성기 만들기
challengeType: 26
dashedName: build-a-sentence-maker
---

# --description--

이 실습에서는 문장 템플릿을 사용해 두 가지 다른 스토리를 만듭니다. 변수에 스토리의 여러 부분을 저장한 다음, 스토리를 콘솔에 출력하세요.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. 다음 변수를 `let`를 사용해 선언하세요:
   
    - `adjective`
    - `noun`
    - `verb`
    - `place`
    - `adjective2`
    - `noun2`

1. 위 변수들에 원하는 문자열 값을 할당하세요.

1. `firstStory` 변수를 선언하세요.
   
1. 다음 스토리 템플릿을 사용해 첫 번째 스토리를 만들고 `firstStory` 변수에 할당하세요: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`;

1. 첫 번째 스토리를 `"First story: [firstStory]"` 메시지를 사용해 콘솔에 출력하세요.

1. `adjective`, `noun`, `verb`, `place`, `adjective2`, `noun2` 변수에 새 값을 할당하세요.

1. `secondStory` 변수를 선언하세요.

1. 같은 템플릿을 사용해 다른 스토리를 만들고 `secondStory` 변수에 할당하세요.

1. 두 번째 스토리를 `"Second story: [secondStory]"` 메시지를 사용해 콘솔에 출력하세요.


# --hints--

`adjective` 변수를 선언하세요.

```js
assert.isNotNull(adjective);
```

`noun` 변수를 선언하세요.

```js
assert.isNotNull(noun);
```

`verb` 변수를 선언하세요.

```js
assert.isNotNull(verb);
```

`place` 변수를 선언하세요.

```js
assert.isNotNull(place);
```

`adjective2` 변수를 선언하세요.

```js
assert.isNotNull(adjective2);
```

`noun2` 변수를 선언하세요.

```js
assert.isNotNull(noun2);
```

`adjective` 변수에 문자열 값을 할당하세요.

```js
assert.isString(adjective);
```

`noun` 변수에 문자열 값을 할당하세요.

```js
assert.isString(noun);
```

`verb` 변수에 문자열 값을 할당하세요.

```js
assert.isString(verb);
```

`place` 변수에 문자열 값을 할당하세요.

```js
assert.isString(place);
```

`adjective2` 변수에 문자열 값을 할당하세요.

```js
assert.isString(adjective2);
```

`noun2` 변수에 문자열 값을 할당하세요.

```js
assert.isString(noun2);
```

`firstStory` 변수를 선언하세요.

```js
assert.isNotNull(firstStory);
```

첫 번째 스토리에 올바른 스토리 형식을 사용하세요: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. 공백에 주의하세요.

```js
const _initialValues = {}
for (const name of ['adjective', 'noun', 'noun2', 'place', 'verb', 'adjective2']) {
    const match = code.match(new RegExp(String.raw`${name}\s*=\s*('|"|${'`'})(?<${name}>.*?)\1(?:,|;|\s|$)`, 'm'));
    _initialValues[name] = match ? match.groups[name] : null;
}

const expected = `Once upon a time, there was a(n) ${_initialValues['adjective']} ${_initialValues['noun']} who loved to eat ${_initialValues['noun2']}. The ${_initialValues['noun']} lived in a ${_initialValues['place']} and had ${_initialValues['adjective2']} nostrils that blew fire when it was ${_initialValues['verb']}.`;

assert.strictEqual(firstStory, expected);
```

선언한 변수를 올바른 순서로 사용해 첫 번째 스토리를 조립하세요.

```js
assert.match(
  __helpers.removeJSComments(code),
  /firstStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

`"First story: [firstStory]"` 메시지를 사용해 첫 번째 스토리를 로그하세요.

```js
const condition1 = /console\.log\(\s*["']First\s+story:\s+["']\s*\+\s*firstStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`First\s+story:\s+\$\{firstStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`secondStory` 변수를 선언하세요.

```js
assert.isNotNull(secondStory);
```

두 번째 스토리를 위해 `adjective` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective\s*=\s*/g), 2);
```

두 번째 스토리를 위해 `noun` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun\s*=\s*/g), 2);
```

두 번째 스토리를 위해 `verb` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/verb\s*=\s*/g), 2);
```

두 번째 스토리를 위해 `place` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/place\s*=\s*/g), 2);
```

두 번째 스토리를 위해 `adjective2` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/adjective2\s*=\s*/g), 2);
```

두 번째 스토리를 위해 `noun2` 변수를 재할당하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/noun2\s*=\s*/g), 2);
```

두 번째 스토리에 올바른 스토리 형식을 사용하세요: `"Once upon a time, there was a(n) [adjective] [noun] who loved to eat [noun2]. The [noun] lived in a [place] and had [adjective2] nostrils that blew fire when it was [verb]."`. 공백에 주의하세요.

```js
const expected = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;
assert.strictEqual(secondStory, expected);
```

선언한 변수를 올바른 순서로 사용해 두 번째 스토리를 조립하세요.

```js
assert.match(
  __helpers.removeJSComments(code),
  /secondStory\s*=\s*.*?adjective.*?noun.*?noun2.*?noun.*?place.*?adjective2.*?verb/
);
```

`"Second story: [secondStory]"` 형식을 사용해 두 번째 스토리를 로그하세요.

```js
const condition1 = /console\.log\(\s*["']Second\s+story:\s+["']\s*\+\s*secondStory\s*\);?/gm.test(code);
const condition2 = /console\.log\(\s*`Second\s+story:\s+\$\{secondStory\}`\s*\);?/gm.test(code);
assert.isTrue(condition1 || condition2);
```

`firstStory`는 `secondStory`와 같으면 안 됩니다.

```js
assert.notEqual(firstStory, secondStory);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
let adjective = "funny";
let noun = "dragon";
let verb = "jumping";
let place = "garden";
let adjective2 = "sparkling";
let noun2 = "cakes";

const firstStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("First story: " + firstStory);

adjective = "cute";
noun = "puppy";
verb = "barking";
place = "park";
adjective2 = "colorful";
noun2 = "flower";

const secondStory = `Once upon a time, there was a(n) ${adjective} ${noun} who loved to eat ${noun2}. The ${noun} lived in a ${place} and had ${adjective2} nostrils that blew fire when it was ${verb}.`;

console.log("Second story: " + secondStory);
```
