---
id: 587d78af367417b2b2512b03
title: アンケートフォームを作成する
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**目的:** <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a> と機能的に似たアプリを作成してください。**このデモプロジェクトをコピーしないでください**。

**ユーザーストーリー:**

1. `h1` 要素に `id` が `title` のページタイトルがあること
1. `p` 要素に `id` が `description` の短い説明があること
1. `form` 要素に `id` が `survey-form` のものがあること
1. form 要素内に、`input` が `id` で `name` が `type` の `text` フィールドに名前を**必ず**入力すること
1. form 要素内に、`input` が `id` の `email` フィールドにメールアドレスを**必ず**入力すること
1. メールアドレスの形式が正しくない場合、HTML5 のバリデーションエラーが表示されること
1. form 内に、`input` が `id` の `number` フィールドに数字を入力できること
1. 数字入力は、数字以外を入力できないか、ブラウザによっては HTML5 のバリデーションエラーが表示されること
1. `min` と `max` 属性で定義された範囲外の数字を入力すると、HTML5 のバリデーションエラーが表示されること
1. 名前、メール、数字の入力フィールドには、それぞれの目的を説明する `label` 要素があり、ID は `id="name-label"`、`id="email-label"`、`id="number-label"` であること
1. 名前、メール、数字の入力フィールドには、それぞれの説明や指示を示すプレースホルダーテキストがあること
1. form 要素内に、`select` が `id` の `dropdown` ドロップダウン要素があり、選択肢が2つ以上あること
1. form 要素内に、`name` 属性でグループ化された2つ以上のラジオボタンの中から選択できること
1. form 要素内に、`value` 属性を持つ複数のチェックボックスから複数選択できること
1. form 要素内に、追加コメント用の `textarea` があること
1. form 要素内に、`id` が `submit` のボタンがあり、すべての入力を提出できること

ユーザーストーリーを満たし、以下のすべてのテストに合格してこのプロジェクトを完成させてください。自分のスタイルを加えてください。楽しくコーディングしましょう！

**注意:** CSS を適用するために、HTML に `<link rel="stylesheet" href="styles.css">` を追加してください。

# --hints--

`h1` 要素に `id` が `title` のものがあるはずです。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

`#title` は空であってはいけません。

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`p` 要素に `id` が `description` のものがあるはずです。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

`#description` は空であってはいけません。

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`form` 要素に `id` が `survey-form` のものがあるはずです。

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

`input` 要素に `id` が `name` のものがあるはずです。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#name` は `type` が `text` のものを持つはずです。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

`#name` は入力必須にしてください。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#name` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

`input` 要素に `id` が `email` のものがあるはずです。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#email` は `type` が `email` のものを持つはずです。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

`#email` は入力必須にしてください。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#email` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

`input` 要素に `id` が `number` のものがあるはずです。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#number` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

`#number` は `type` が `number` のものを持つはずです。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

`#number` は数値の値を持つ `min` 属性を持つはずです。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

`#number` は数値の値を持つ `max` 属性を持つはずです。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

`label` 要素に `id` が `name-label` のものがあるはずです。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 要素に `id` が `email-label` のものがあるはずです。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 要素に `id` が `number-label` のものがあるはずです。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`#name-label` には入力を説明するテキストが含まれているはずです。

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#email-label` には入力を説明するテキストが含まれているはずです。

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#number-label` には入力を説明するテキストが含まれているはずです。

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#name-label` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

`#email-label` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

`#number-label` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

`#name` は `placeholder` 属性と値を持つはずです。

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#email` は `placeholder` 属性と値を持つはずです。

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#number` は `placeholder` 属性と値を持つはずです。

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`select` フィールドには `id` が `dropdown` のものがあるはずです。

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

`#dropdown` には少なくとも2つの選択可能（無効化されていない）な `option` 要素があるはずです。

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

`#dropdown` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

`input` が少なくとも2つあり、`type` が `radio` のもの（ラジオボタン）があるはずです。

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

少なくとも2つのラジオボタンが `#survey-form` の子孫要素であるはずです。

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

すべてのラジオボタンは `value` 属性と値を持つはずです。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

すべてのラジオボタンは `name` 属性と値を持つはずです。

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

すべてのラジオボタングループには少なくとも2つのラジオボタンがあるはずです。

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

`input` が少なくとも2つあり、`type` が `checkbox` のもの（チェックボックス）が `#survey-form` の子孫要素であるはずです。

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 内のすべてのチェックボックスは `value` 属性と値を持つはずです。

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

`textarea` が少なくとも1つあり、`#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

`input` または `button` 要素に `id` が `submit` のものがあるはずです。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

`#submit` は `type` が `submit` のものを持つはずです。

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

`#submit` は `#survey-form` の子孫要素であるはずです。

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
