---
id: 587d78af367417b2b2512b03
title: Створіть форму опитування
challengeType: 14
saveSubmissionToDB: true
forumTopicId: 301145
dashedName: build-a-survey-form
---

# --description--

**Мета:** Створіть застосунок, який функціонально схожий на <a href="https://survey-form.freecodecamp.rocks" target="_blank" rel="noopener noreferrer nofollow">https://survey-form.freecodecamp.rocks</a>. **Не копіюйте цей демонстраційний проєкт**.

**Історія користувача:**

1. Ви повинні мати заголовок сторінки в елементі `h1` з ідентифікатором `id` рівним `title`
1. Ви повинні мати короткий опис у елементі `p` з ідентифікатором `id` рівним `description`
1. Ви повинні мати елемент `form` з ідентифікатором `id` рівним `survey-form`
1. Всередині елемента форми ви **повинні** ввести своє ім’я у полі введення `input`, яке має атрибут ідентифікатора `id` зі значенням `name` та атрибут `type` зі значенням `text`
1. Всередині елемента форми ви **повинні** ввести свою е-пошту у полі введення `input`, яке має атрибут ідентифікатора `id` зі значенням `email`
1. Якщо ви введете е-пошту у неправильному форматі, ви побачите помилку валідації HTML5
1. Всередині форми ви можете ввести число у полі введення `input`, яке має атрибут ідентифікатора `id` зі значенням `number`
1. Поле введення числа не повинно приймати нечислові символи, або забороняючи їх введення, або показуючи помилку валідації HTML5 (залежно від браузера).
1. Якщо ви введете числа поза межами діапазону поля введення числа, визначеного атрибутами `min` та `max`, ви побачите помилку валідації HTML5
1. Для полів введення імені, е-пошти та числа ви повинні бачити відповідні елементи `label` у формі, які описують призначення кожного поля з такими ідентифікаторами: `id="name-label"`, `id="email-label"` та `id="number-label"`
1. Для полів введення імені, е-пошти та числа ви повинні бачити текст заповнювача, який дає опис або інструкції для кожного поля
1. Всередині елемента форми ви повинні мати випадний елемент `select` з ідентифікатором `id`, рівним `dropdown`, і принаймні двома варіантами для вибору
1. Всередині елемента форми ви можете вибрати варіант із групи принаймні двох радіокнопок, згрупованих за допомогою атрибута `name`
1. Всередині елемента форми ви можете вибрати кілька полів із серії прапорців, кожен з яких повинен мати атрибут `value`
1. Всередині елемента форми вам пропонується поле `textarea` для додаткових коментарів
1. Всередині елемента форми вам пропонується кнопка з атрибутом `id` рівним `submit` для відправлення всіх введених даних

Виконайте історії користувача та пройдіть усі тести нижче, щоб завершити цей проєкт. Додайте свій особистий стиль. Успішного кодування!

**Примітка:** Обов’язково додайте `<link rel="stylesheet" href="styles.css">` у свій HTML, щоб підключити таблицю стилів і застосувати CSS

# --hints--

Ви повинні мати елемент `h1` з атрибутом ідентифікатора `id`, рівним `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Ваш `#title` не повинен бути порожнім.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ви повинні мати елемент `p` з атрибутом ідентифікатора `id`, рівним `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Ваш `#description` не повинен бути порожнім.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ви повинні мати елемент `form` з атрибутом ідентифікатора `id`, рівним `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Ви повинні мати елемент `input` з атрибутом ідентифікатора `id`, рівним `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ваш `#name` повинен мати атрибут ідентифікатора `type`, рівний `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Ваш `#name` повинен вимагати введення.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ваш `#name` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Ви повинні мати елемент `input` з атрибутом ідентифікатора `id`, рівним `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ваш `#email` повинен мати атрибут ідентифікатора `type`, рівний `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Ваш `#email` повинен вимагати введення.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Ваш `#email` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Ви повинні мати елемент `input` з атрибутом ідентифікатора `id`, рівним `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Ваш `#number` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Ваш `#number` повинен мати атрибут ідентифікатора `type`, рівний `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Ваш `#number` повинен мати атрибут `min` з числовим значенням.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Ваш `#number` повинен мати атрибут `max` з числовим значенням.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Ви повинні мати елемент `label` з атрибутом ідентифікатора `id`, рівним `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ви повинні мати елемент `label` з атрибутом ідентифікатора `id`, рівним `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ви повинні мати елемент `label` з атрибутом ідентифікатора `id`, рівним `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Ваш `#name-label` повинен містити текст, який описує вхідні дані.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ваш `#email-label` повинен містити текст, який описує вхідні дані.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ваш `#number-label` повинен містити текст, який описує вхідні дані.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Ваш `#name-label` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Ваш `#email-label` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Ваш `#number-label` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Ваш `#name` повинен мати атрибут і значення `placeholder`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ваш `#email` повинен мати атрибут і значення `placeholder`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ваш `#number` повинен мати атрибут і значення `placeholder`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Ви повинні мати поле `select` з атрибутом ідентифікатора `id`, рівним `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Ваш `#dropdown` повинен мати принаймні два вибірні (не вимкнені) елементи `option`.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Ваш `#dropdown` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Ви повинні мати принаймні два елементи `input` з атрибутом ідентифікатора `type`, рівним `radio` (радіокнопки).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Ви повинні мати принаймні дві радіокнопки, які є нащадками `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Усі ваші радіокнопки повинні мати атрибут і значення `value`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Усі ваші радіокнопки повинні мати атрибут і значення `name`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Кожна група радіокнопок повинна мати принаймні 2 радіокнопки.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
});

assert.isAbove(groupKeys.length, 0);
```

Ви повинні мати принаймні два елементи `input` з атрибутом ідентифікатора `type`, рівним `checkbox` (прапорці), які є нащадками `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Усі ваші прапорці всередині `#survey-form` повинні мати атрибут і значення `value`.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Ви повинні мати принаймні один елемент `textarea`, який є нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Ви повинні мати елемент `input` або `button` з атрибутом ідентифікатора `id`, рівним `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Ваш `#submit` повинен мати атрибут ідентифікатора `type`, рівний `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Ваш `#submit` повинен бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html

```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
    <title>Survey Form</title>
  </head>
  <body>
    <h1>Survey Form</h1>
    <p>The card below was built as a sample survey form for freeCodeCamp.</p>
    <main id="main">
      <h1 id="title">Join the Togepi Fan Club!</h1>
      <p id="description">
        Enter your information here to receive updates about club activities,
        our monthly newsletter, and other email communications.
      </p>
      <form id="survey-form" action="#">
        <label for="name" id="name-label"
          ><p>Name:</p>
          <input type="text" id="name" placeholder="e.g. John Smith" required />
        </label>
        <label for="email" id="email-label"
          ><p>Email:</p>
          <input
            type="email"
            id="email"
            placeholder="e.g. john.smith@email.com"
            required
          />
        </label>
        <label for="age" id="number-label"
          ><p>Age<em>(optional)</em>:</p>
          <input
            type="number"
            id="number"
            placeholder="e.g. 19"
            min="18"
            max="99"
          />
        </label>
        <label for="interest" id="interest-label"
          ><p>What are you most interested in?</p>
          <select id="dropdown">
            <option selected disabled hidden></option>
            <option id="battles">Battling</option>
            <option id="breeding">Breeding</option>
            <option id="catching">Completing my Pokedex</option>
            <option id="exploring">Exploring new regions</option>
          </select>
        </label>
        <p>Who is your favourite Pokemon?</p>
        <label for="togepi">
          <input
            id="togepi"
            type="radio"
            name="favorite"
            value="togepi"
          />Togepi!
        </label>
        <label for="pikachu">
          <input
            id="pikachu"
            type="radio"
            name="favorite"
            value="pikachu"
          />Pikachu
        </label>
        <label for="other">
          <input id="other" type="radio" name="favorite" value="other" />Other
        </label>
        <p>Which communications do you want to receive?</p>
        <label for="newsletter">
          <input
            id="newsletter"
            type="checkbox"
            name="communications"
            value="newsletter"
          />Newsletter
        </label>
        <label for="events">
          <input
            id="events"
            type="checkbox"
            name="communications"
            value="events"
          />Event updates
        </label>
        <label for="updates">
          <input
            id="updates"
            type="checkbox"
            name="communications"
            value="updates"
          />Club updates
        </label>
        <p>Any other information you would like to share?</p>
        <textarea id="additional-information" rows="4" cols="50">
You can provide comments, suggestions, or feedback here.</textarea
        >
        <p>
          <em
            >Please note: This form is a proof of concept. Submitting the form
            will not actually transmit your data.</em
          >
        </p>
        <button type="Submit" id="submit">Submit</button>
      </form>
    </main>
  </body>
  <footer>
    <a href="../">Return to Project List</a> |
    <a href="https://www.nhcarrigan.com">Return to HomePage</a>
  </footer>
</html>
```

```css
main {
  text-align: center;
  background-color: #92869c;
  background-blend-mode: lighten;
  max-width: 500px;
  margin: 20px auto;
  border-radius: 50px;
  box-shadow: 10px 10px rgba(0, 0, 0, 0.5);
  color: black;
}
body {
  text-align: center;
  background: #3a3240;
  color: white;
}
input,
textarea,
select,
button {
  background: #3a3240;
  color: white;
}
a {
  color: white;
}
```
