---
id: 69899123889dcf5b19be1847
title: реалізувати алгоритм пошуку у глибину
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

У цій лабораторній роботі ви реалізуєте алгоритм обходу графа, який називається <dfn>пошук у глибину</dfn> (DFS).

Якщо <dfn>пошук у ширину</dfn> (BFS) досліджує сусідів рівень за рівнем, то <dfn>пошук у глибину</dfn> занурюється якомога глибше по одному шляху ребер, перш ніж повернутися назад.

Коли алгоритм досягає кінця шляху (вузол без невідвіданих сусідів), він повертається до останнього вузла, у якого ще є не досліджені ребра, і продовжує пошук звідти.

Щоб реалізувати цей алгоритм ітеративно, вам знадобиться **стек**. У JavaScript стандартний `Array` ідеально підходить для стека, використовуючи методи `.push()` та `.pop()`, дотримуючись принципу **останній зайшов — перший вийшов (LIFO)**. Це гарантує, що сусід, який був доданий у стек останнім, буде наступним для дослідження.

Простим вихідним результатом цього алгоритму є список вузлів, до яких можна дістатися з заданого вузла. Тому вам також потрібно відстежувати вузли, які ви відвідуєте.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з іменем `dfs`.
1. Функція `dfs` повинна приймати два аргументи:

    - `graph`: матриця суміжності (масив масивів).
    - `root`: числова мітка вузла (початковий індекс), яка є числовим значенням вузла між `0` та `n - 1`, де `n` — загальна кількість вузлів у графі.

1. Функція повинна повертати масив, що містить усі мітки вузлів, до яких можна дістатися з початкового вузла.

# --hints--

Ви повинні мати функцію з іменем `dfs`, яка приймає два аргументи.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` має повертати список з `1`, `2`, `3` та `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` має повертати список з `1`, `2`, `3` та `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` має повертати `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` має повертати список з `3` та `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` має повертати список з `0` та `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
