---
id: 60f02e7361b68405e27b62a5
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kupamba fomu kwa kubuni fomu ya usajili.

Muundo wa msingi wote wa HTML (`DOCTYPE`, `html`, `head`, na `body`) umewekwa tayari kwa ajili yako.

Ndani ya `body`, ongeza kichwa cha sehemu ili kutoa muktadha kwa fomu kwa kutumia kipengele cha `h1` chenye maandishi `Registration Form`.

# --hints--

Unapaswa kuongeza `h1` ndani ya `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Unapaswa kumpa `h1` maandishi `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
