---
id: 60f030d388cb74067cf291c3
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya kichwa cha sehemu, tumia maandishi yafuatayo ndani ya kipengele cha aya ili kuhamasisha watumizi kujiandikisha:

```md
Please fill out this form with the required information
```

# --hints--

Unapaswa kuongeza kipengele cha `p` ndani ya `body`.

```js
assert.exists(document.querySelector('body > p'));
```

Unapaswa kuongeza kipengele cha `p` chini ya `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

Unapaswa kumpa kipengele cha `p` maandishi ya `Please fill out this form with the required information`.

```js
assert.equal(document.querySelector('p')?.innerText, 'Please fill out this form with the required information');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
--fcc-editable-region--
<body>
  <h1>Registration Form</h1>
  
</body>
--fcc-editable-region--
</html>
```
