---
id: 60f80e0081e0f2052ae5b505
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Kuweka sifa ya kipengele `type` ya kipengele cha `input` ni muhimu ili kivinjari kijue aina gani ya data kinapaswa kutarajia. Ikiwa `type` haijaainishwa, kivinjari kita chagua `text` kama chaguo-msingi.

Toa sifa ya kipengele `input` yenye thamani ya `type` kwa vipengele vya kwanza viwili vya `text`, sifa ya kipengele `type` yenye thamani ya `email` kwa kipengele cha tatu, na sifa ya kipengele `type` yenye thamani ya `password` kwa kipengele cha nne.

Aina ya `email` inaruhusu barua pepe tu zenye `@` na `.` katika kikoa.
Aina ya `password` huficha ingizo, na hutoa onyo ikiwa tovuti haitumii HTTPS.

# --hints--

Unapaswa kutoa sifa ya kipengele `input` yenye thamani ya `type` kwa kipengele cha kwanza cha `text`.

```js
assert.equal(document.querySelector('input')?.type, 'text');
```

Unapaswa kutoa sifa ya kipengele `input` yenye thamani ya `type` kwa kipengele cha pili cha `text`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'text');
```

Unapaswa kutoa sifa ya kipengele `input` yenye thamani ya `type` kwa kipengele cha tatu cha `email`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'email');
```

Unapaswa kutoa sifa ya kipengele `input` yenye thamani ya `type` kwa kipengele cha nne cha `password`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.type, 'password');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" /></label>
      <label for="email">Enter Your Email: <input id="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
