---
id: 60f805f813eaf2049bc2ceea
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Panga ndani kipengele cha `input` ndani ya kila `label`. Hakikisha unaongeza kila `input` baada ya maandishi ya `label`, na jumuisha nafasi baada ya nukta mbili.

# --hints--

Unapaswa kuongeza vipengele vinne vya `input` kwenye kipengele cha `fieldset`.

```js
assert.equal(document.querySelectorAll('fieldset input')?.length, 4);
```

Unapaswa kupanga ndani vipengele vya `input` ndani ya vipengele vya `label`.

```js
assert.equal(document.querySelectorAll('label input')?.length, 4);
```

Unapaswa kuongeza `input` ya kwanza baada ya maandishi ya `label` `Enter Your First Name:`, na jumuisha nafasi baada ya nukta mbili.

```js
const query = /^Enter Your First Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[0]?.innerHTML.trim(), query);
```

Unapaswa kuongeza `input` ya pili baada ya maandishi ya `label` `Enter Your Last Name:`, na jumuisha nafasi baada ya nukta mbili.

```js
const query = /^Enter Your Last Name:\s+<input/
assert.match(document.querySelectorAll('label')?.[1]?.innerHTML.trim(), query);
```

Unapaswa kuongeza `input` ya tatu baada ya maandishi ya `label` `Enter Your Email:`, na jumuisha nafasi baada ya nukta mbili.

```js
const query = /^Enter Your Email:\s+<input/
assert.match(document.querySelectorAll('label')?.[2]?.innerHTML.trim(), query);
```

Unapaswa kuongeza `input` ya nne baada ya maandishi ya `label` `Create a New Password:`, na jumuisha nafasi baada ya nukta mbili.

```js
const query = /^Create a New Password:\s+<input/
assert.match(document.querySelectorAll('label')?.[3]?.innerHTML.trim(), query);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label>Enter Your First Name:</label>
      <label>Enter Your Last Name:</label>
      <label>Enter Your Email:</label>
      <label>Create a New Password:</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
