---
id: 60fab4a123ce4b04526b082b
title: Hatua 24
challengeType: 0
dashedName: step-24
---

# --description--

Kwa sasa watumizi wanaweza kuwasilisha fomu bila kukagua ingizo za redio. Ingawa hapo awali ulitumia sifa ya kipengele `required` kuonyesha kuwa ingizo ni la lazima, haitafanya kazi katika kesi hii kwa sababu kuongeza `required` kwa ingizo zote mbili kutawasilisha taarifa isiyo sahihi kwa watumizi.

Ili kutatua hili, unaweza kutoa muktadha wa kinachohitajika kwa kuongeza kipengele cha `legend` chenye maandishi `Account type (required)` kabla ya vipengele vya `label` ndani ya `fieldset` ya pili. Kisha ongeza sifa ya kipengele `checked` kwa ingizo la `Personal` kuhakikisha fomu inawasilishwa ikiwa na data inayohitajika ndani yake.

# --hints--

Vipengele vyako vya `input` vinapaswa kubaki `type` ya `radio`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-child(2) input[type="radio"]'), 2);
```

Unapaswa kuongeza kipengele cha `legend` ndani ya kipengele cha `fieldset` cha pili.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) > legend'), 1);
```

Unapaswa kuongeza maandishi ya `Account type (required)` kwenye kipengele cha `legend`.

```js
assert.equal(document.querySelector('fieldset:nth-of-type(2) > legend')?.textContent, 'Account type (required)');
```

Unapaswa kutoa sifa ya kipengele `checked` kwa ingizo la `Personal`.

```js
assert.isTrue(document.querySelector('fieldset:nth-child(2) input[type="radio"]')?.hasAttribute('checked'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      
      <label><input type="radio" name="account-type" value="personal" /> Personal</label>
--fcc-editable-region--
      <label><input type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
