---
id: 587d78af367417b2b2512b03
title: Створіть форму для опитування
challengeType: 14
dashedName: build-a-survey-form
demoType: onClick
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Ви повинні мати назву сторінки в елементі `h1` з `id` зі значенням `title`.
1. Ви повинні мати коротке пояснення в елементі `p` з `id` зі значенням `description`.
1. Ви повинні мати елемент `form` з `id` зі значенням `survey-form`.
1. Inside the form element you should have a required `input` field to enter your name that has an `id` of `name` and a `type` of `text`.
1. Inside the form element you should have a required `input` field to enter your email that has an `id` of `email`.
1. Якщо ви введете неправильно сформовану електронну адресу, то повинні побачити помилку перевірки HTML5.
1. Inside the form element you should have an `input` field to enter a number that has an `id` of `number`.
1. Ввід для чисел повинен приймати лише числа, або забороняючи вводити їх, або показуючи помилку перевірки HTML5 (залежно від вашого браузера).
1. Якщо ви введете число поза межами діапазону чисел, що визначено атрибутами `min` та `max`, ви повинні побачити помилку перевірки HTML5.
1. Для імені, електронної пошти та полів введення чисел ви повинні мати відповідні елементи `label` в формі, що описують мету кожного поля відповідними ідентифікаторами: `name-label`, `email-label` та `number-label`.
1. Для імені, електронної пошти та полів введення чисел ви повинні мати текст заповнювача з описом або інструкцією до кожного поля.
1. В елементі форми ви повинні мати випадний елемент `select` з `id` зі значенням `dropdown` та принаймні два варіанти вибору.
1. В елементі форми ви можете вибрати опцію з групи принаймні двох радіокнопок, які згруповані з використанням атрибуту `name`.
1. В елементі форми ви можете вибрати декілька полів з серії прапорців, кожен з яких повинен мати атрибут `value`.
1. В елементі форми ви повинні мати `textarea` для додаткових коментарів.
1. В елементі форми ви повинні мати кнопку з `id` зі значенням `submit` для відправки всіх введень.

# --hints--

Ви повинні мати елемент `h1` з `id` зі значенням `title`.

```js
const el = document.getElementById('title');
assert(!!el && el.tagName === 'H1');
```

`#title` не повинен бути порожнім.

```js
const el = document.getElementById('title');
assert(!!el && el.innerText.length > 0);
```

Ви повинні мати елемент `p` з `id` зі значенням `description`.

```js
const el = document.getElementById('description');
assert(!!el && el.tagName === 'P');
```

`#description` не повинен бути порожнім.

```js
const el = document.getElementById('description');
assert(!!el && el.innerText.length > 0);
```

Ви повинні мати елемент `form` з `id` зі значенням `survey-form`.

```js
const el = document.getElementById('survey-form');
assert(!!el && el.tagName === 'FORM');
```

Ви повинні мати елемент `input` з `id` зі значенням `name`.

```js
const el = document.getElementById('name');
assert(!!el && el.tagName === 'INPUT');
```

`#name` повинен мати `type` зі значенням `text`.

```js
const el = document.getElementById('name');
assert(!!el && el.type === 'text');
```

`#name` повинен вимагати введення.

```js
const el = document.getElementById('name');
assert(!!el && el.required);
```

`#name` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert(!!el);
```

Ви повинні мати елемент `input` з `id` зі значенням `email`.

```js
const el = document.getElementById('email');
assert(!!el && el.tagName === 'INPUT');
```

`#email` повинен мати `type` зі значенням `email`.

```js
const el = document.getElementById('email');
assert(!!el && el.type === 'email');
```

`#email` повинен вимагати введення.

```js
const el = document.getElementById('email');
assert(!!el && el.required);
```

`#email` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert(!!el);
```

Ви повинні мати елемент `input` з `id` зі значенням `number`.

```js
const el = document.getElementById('number');
assert(!!el && el.tagName === 'INPUT');
```

`#number` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert(!!el);
```

`#number` повинен мати `type` зі значенням `number`.

```js
const el = document.getElementById('number');
assert(!!el && el.type === 'number');
```

`#number` повинен мати атрибут `min` з числовим значенням.

```js
const el = document.getElementById('number');
assert(!!el && el.min && isFinite(el.min));
```

`#number` повинен мати атрибут `max` з числовим значенням.

```js
const el = document.getElementById('number');
assert(!!el && el.max && isFinite(el.max));
```

Ви повинні мати елемент `label` з `id` зі значенням `name-label`.

```js
const el = document.getElementById('name-label');
assert(!!el && el.tagName === 'LABEL');
```

Ви повинні мати елемент `label` з `id` зі значенням `email-label`.

```js
const el = document.getElementById('email-label');
assert(!!el && el.tagName === 'LABEL');
```

Ви повинні мати елемент `label` з `id` зі значенням `number-label`.

```js
const el = document.getElementById('number-label');
assert(!!el && el.tagName === 'LABEL');
```

`#name-label` має містити текст, який описує введені дані.

```js
const el = document.getElementById('name-label');
assert(!!el && el.innerText.length > 0);
```

`#email-label` має містити текст, який описує введені дані.

```js
const el = document.getElementById('email-label');
assert(!!el && el.innerText.length > 0);
```

`#number-label` має містити текст, який описує введені дані.

```js
const el = document.getElementById('number-label');
assert(!!el && el.innerText.length > 0);
```

`#name-label` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert(!!el);
```

`#email-label` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert(!!el);
```

`#number-label` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert(!!el);
```

Your `#name` should have a `placeholder` attribute and value.

```js
const el = document.getElementById('name');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

Your `#email` should have a `placeholder` attribute and value.

```js
const el = document.getElementById('email');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

Your `#number` should have a `placeholder` attribute and value.

```js
const el = document.getElementById('number');
assert(!!el && !!el.placeholder && el.placeholder.length > 0);
```

You should have a `select` field with an `id` of `dropdown`.

```js
const el = document.getElementById('dropdown');
assert(!!el && el.tagName === 'SELECT');
```

`#dropdown` повинен мати щонайменше два елементи `option` з можливістю вибору (не відключені).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert(els.length >= 2);
```

`#dropdown` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert(!!el);
```

Ви повинні мати щонайменше два елементи `input` з `type` зі значенням `radio` (радіокнопки).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert(els.length >= 2);
```

Ви повинні мати щонайменше дві радіокнопки, які є нащадками `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert(els.length >= 2);
```

Всі радіокнопки повинні мати атрибут та значення `value`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][value=""], input[type="radio"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

Всі радіокнопки повинні мати атрибут та значення `name`.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll('input[type="radio"][name=""], input[type="radio"]:not([name])');
assert(els1.length > 0 && els2.length === 0);
```

В кожній групі радіокнопок мають бути щонайменше 2 радіокнопки.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  })
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert(false);
})

assert(groupKeys.length > 0);
```

Ви повинні мати щонайменше два елементи `input` з `type` зі значенням `checkbox` (прапорцями), що є нащадками `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert(els.length >= 2);
```

Всі прапорці всередині `#survey-form` повинні мати атрибут та значення `value`.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll('#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])');
assert(els1.length > 0 && els2.length === 0);
```

Ви повинні мати щонайменше один елемент `textarea`, що є нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert(!!el);
```

Ви повинні мати елемент `input` або `button` з `id` зі значенням `submit`.

```js
const el = document.getElementById('submit');
assert(!!el && (el.tagName === 'INPUT' || el.tagName === 'BUTTON'));
```

`#submit` повинен мати `type` зі значенням `submit`.

```js
const el = document.getElementById('submit');
assert(!!el && el.type === 'submit');
```

`#submit` має бути нащадком `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert(!!el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
</head>

<body>

</body>

</html> 
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
