---
id: 66cc14ca14e65e7a9ebecb80
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이제 계산기가 두 숫자를 더할 수 있으니, 두 숫자를 빼는 함수를 만들어 보세요.

`calculateDifference`라는 이름의 함수를 선언하고, 두 개의 매개변수 `num1`과 `num2`를 받게 하세요.

함수 안에서 두 매개변수의 차이를 반환하세요.

# --hints--

`calculateDifference`라는 함수가 있어야 합니다.

```js
assert.isFunction(calculateDifference);
```

`calculateDifference` 함수는 `num1`와 `num2`라는 두 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateDifference` 함수는 두 매개변수의 차이를 반환해야 합니다.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
