---
id: 66cc1a3a39aef47d6473cb2f
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

다음 단계는 계산기에 나누기 기능을 추가하는 것입니다.

두 개의 매개변수 `calculateQuotient`과 `num1`를 받는 `num2`라는 함수를 선언하세요.

함수 내부에서 두 매개변수의 몫을 반환하세요.

그다음, `console.log`와 `calculateQuotient` 인수를 사용하여 `7` 함수를 호출하는 `11`를 추가하세요.

# --hints--

`calculateQuotient`라는 함수를 선언하세요.

```js
assert.isFunction(calculateQuotient);
```

`calculateQuotient` 함수는 `num1`와 `num2`라는 두 개의 매개변수를 가져야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

`calculateQuotient` 함수는 두 매개변수의 몫을 반환해야 합니다.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

`console.log`가 `calculateQuotient`와 `7` 인수를 사용하여 `11` 함수를 호출해야 합니다.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
