---
id: 66cc0a9e06e00b75d6782be9
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이전 수업에서 함수 호출 방법을 배웠습니다. 함수 호출은 함수 내부의 코드를 실행하는 것을 의미합니다.

함수를 호출하는 방법을 다시 알려드립니다:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

`console.log` 문장을 추가하고, 그 안에서 `addTwoAndSeven` 함수를 호출하세요.

이제 `2`와 `7`의 합이 콘솔에 출력되어야 합니다.

# --hints--

코드에 `console.log`가 있어야 합니다.

```js
assert.match(code, /console\.log(.*)/);
```

`addTwoAndSeven` 안에서 `console.log` 함수를 호출해야 합니다.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
