---
id: 66cc16abc30b367ba1d1f9ef
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 `calculateDifference` 함수를 테스트할 시간입니다.

먼저 `console.log`를 추가하여 `calculateDifference` 함수를 인수 `22`와 `5`로 호출하세요.

그다음, 두 번째 `console.log`를 추가하여 `calculateDifference` 함수를 인수 `12`와 `1`로 호출하세요.

마지막으로, 세 번째 `console.log`를 추가하여 `calculateDifference` 함수를 인수 `17`와 `9`로 호출하세요.

# --hints--

`console.log`가 `calculateDifference` 함수를 인수 `22`와 `5`로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

`console.log`가 `calculateDifference` 함수를 인수 `12`와 `1`로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

`console.log`가 `calculateDifference` 함수를 인수 `17`와 `9`로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
