---
id: 66cc0de2920bec775f610424
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

함수들이 예상대로 작동하더라도, 코드에 반복이 많습니다.

계산기를 만들고 있으니, 더하는 모든 가능한 숫자 조합마다 함수를 만들고 싶지 않을 것입니다.

두 숫자를 더할 수 있는 하나의 재사용 가능한 함수를 만드는 것이 더 좋습니다.

지금까지 작성한 모든 코드를 삭제하세요.

다음 단계에서는 함수에서 매개변수와 인수를 다루는 방법을 검토할 것입니다.

# --hints--

`addTwoAndSeven`라는 함수가 없어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

`addTwoAndSeven`에 대한 함수 호출이 없어야 합니다.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

`addThreeAndFour`라는 함수가 없어야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

`addThreeAndFour`에 대한 함수 호출이 없어야 합니다.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
