---
id: 66cc0f1ae40802781b2ea972
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이전 수업에서 함수 매개변수를 다루는 방법을 배웠습니다.

함수 매개변수는 함수 선언에서 정의된 변수로, 자리 표시자 역할을 합니다.

다음은 매개변수를 가진 함수의 예입니다:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

두 개의 매개변수 `calculateSum`과 `num1`를 받는 `num2`라는 함수를 선언하세요.

# --hints--

`calculateSum`라는 함수를 선언해야 합니다.

```js
assert.isFunction(calculateSum);
```

`calculateSum` 함수는 두 개의 매개변수 `num1`와 `num2`를 받아야 합니다.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
