---
id: 66cc2739f687e484d50bb6f1
title: 18단계
challengeType: 1
dashedName: step-18
---

# --description--

계산기의 마지막 기능은 제곱근 기능입니다.

제곱은 어떤 수를 자기 자신과 곱한 수입니다. 따라서 제곱근은 자기 자신과 곱했을 때 원래 수가 되는 수입니다.

예를 들어, `9`의 제곱근은 `3`입니다. 왜냐하면 `3 * 3 = 9`이기 때문입니다.

JavaScript에서 숫자의 제곱근을 구하려면 <dfn>Math.sqrt()</dfn> 메서드를 사용할 수 있습니다.

예시는 다음과 같습니다:

```js
// result: 3
Math.sqrt(9);
```

`calculateSquareRoot`라는 이름의 함수를 선언하세요. 이 함수는 `num` 매개변수를 받고 `num`의 제곱근을 반환해야 합니다.

# --hints--

`calculateSquareRoot`라는 함수가 있어야 합니다.

```js
assert.isFunction(calculateSquareRoot);
```

`calculateSquareRoot` 함수는 `num` 매개변수를 받아야 합니다.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateSquareRoot` 함수는 `num`의 제곱근을 반환해야 합니다.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
