---
id: 66cc21d23238dc8240a8a182
title: 16단계
challengeType: 1
dashedName: step-16
---

# --description--

계산기가 숫자의 제곱을 계산할 수 있으면 좋겠습니다. 숫자의 제곱은 그 숫자에 자기 자신을 곱한 값입니다.

JavaScript에서 숫자의 제곱을 계산하려면 <dfn>Math.pow()</dfn> 메서드를 사용할 수 있습니다. `Math.pow()` 메서드는 두 개의 인수를 받는데, 밑(base) 숫자와 지수(exponent)입니다.

또한 지수 연산자(`**`)를 사용해 숫자의 제곱을 계산할 수도 있습니다.

예시는 다음과 같습니다:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

`calculateSquare`라는 이름의 함수를 선언하고, `num` 매개변수를 받아 `num`의 제곱을 반환하세요.

# --hints--

`calculateSquare`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(calculateSquare);
```

`calculateSquare` 함수는 `num` 매개변수를 받아야 합니다.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateSquare` 함수는 `num`의 제곱을 반환해야 합니다.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
