---
id: 66cc172af7159a7c67804544
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

이제 계산기에 곱셈 기능을 추가할 차례입니다.

두 개의 매개변수 `calculateProduct`과 `num1`를 받는 `num2`라는 함수를 선언하세요.

함수 내부에서 두 매개변수의 곱을 반환하세요.

그다음, `console.log`와 `calculateProduct` 인수를 사용해 `13` 함수를 호출하는 `5`를 추가하세요.

# --hints--

`calculateProduct`라는 함수가 있어야 합니다.

```js
assert.isFunction(calculateProduct);
```

`calculateProduct` 함수는 `num1`와 `num2`라는 두 매개변수를 가져야 합니다.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

`calculateProduct` 함수는 두 매개변수의 곱을 반환해야 합니다.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

`console.log`가 `calculateProduct` 함수를 `13`와 `5` 인수로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
