---
id: 66cc095dfe1682753d2ab030
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이전 수업에서 배운 개념 중 하나는 함수에서 값을 반환하는 방법이었습니다.

함수에서 값을 반환하는 방법을 다시 알려드립니다:

```js
function myFunction() {
  return "Hello World";
}
```

`addTwoAndSeven` 함수 안에서 `2`과 `7`의 합을 반환하세요.

# --hints--

`addTwoAndSeven` 함수는 `2`와 `7`의 합을 반환해야 합니다.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

`addTwoAndSeven` 함수는 `+` 연산자를 사용해 합을 반환해야 합니다.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
