---
id: 66cc10958d72e57907f898b2
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

`calculateSum` 함수 안에서 두 매개변수의 합을 반환해야 합니다.

# --hints--

`calculateSum` 함수는 두 매개변수의 합을 반환해야 합니다.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
