---
id: 66cc06012e46aa6bc9b8c001
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 계산기 앱을 만들면서 함수 작업 방법을 검토합니다.

이전 수업에서 다음과 같이 함수를 선언하는 방법을 배웠습니다:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

먼저 `addTwoAndSeven`라는 함수를 만드세요.

일반 함수 구문이나 화살표 함수 구문 중 하나를 선택할 수 있습니다.

# --hints--

`addTwoAndSeven`라는 함수를 선언해야 합니다.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
