---
id: 66cc0ba5881acb7692cfc4de
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 또 다른 함수를 추가할 시간입니다.

`addThreeAndFour`라는 함수를 선언하고 `3`과 `4`의 합을 반환하세요.

그다음 `addThreeAndFour` 안에서 `console.log` 함수를 호출하여 결과를 확인하세요.

# --hints--

`addThreeAndFour`라는 함수를 선언해야 합니다.

```js
assert.isFunction(addThreeAndFour);
```

`addThreeAndFour` 함수는 `3`과 `4`의 합을 반환해야 합니다.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

`addThreeAndFour` 함수는 `+` 연산자를 사용해 합을 반환해야 합니다.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

`addThreeAndFour` 함수를 `console.log` 안에서 호출해야 합니다.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
