---
id: 66cc26907ff6908402af0149
title: 17단계
challengeType: 1
dashedName: step-17
---

# --description--

`calculateSquare` 함수를 테스트하려면, 함수를 몇 번 호출해야 합니다.

`console.log`를 추가하여 `calculateSquare` 함수를 `2` 인수로 호출하세요.

그다음, `console.log`를 추가하여 `calculateSquare` 함수를 `9` 인수로 호출하세요.

# --hints--

`console.log`가 `calculateSquare` 함수를 `2` 인수로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

`console.log`가 `calculateSquare` 함수를 `9` 인수로 호출해야 합니다.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
