---
id: 66cc0ba5881acb7692cfc4de
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 또 다른 함수를 추가할 시간입니다.

`addThreeAndFour`라는 함수를 선언하고 `3`과 `4`의 합을 반환하세요.

그다음 `addThreeAndFour` 안에서 `console.log` 함수를 호출하여 결과를 확인하세요.

# --hints--

`addThreeAndFour`라는 함수를 선언해야 합니다.

```js
assert.isFunction(addThreeAndFour);
```

`addThreeAndFour` 함수는 `3`과 `4`의 합을 반환해야 합니다.

```js
assert.equal(addThreeAndFour(), 7);
```

`addThreeAndFour` 함수는 `+` 연산자를 사용해 합을 반환해야 합니다.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

`addThreeAndFour` 함수를 `console.log` 안에서 호출해야 합니다.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
