---
id: 66cc12fa504b0479dac479a0
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

함수의 장점은 서로 다른 인수로 호출할 수 있어서 같은 코드를 다양한 값으로 재사용할 수 있다는 점입니다.

다음은 서로 다른 문자열 인수로 함수가 호출되는 예시입니다:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

인수 `console.log`와 `calculateSum`를 사용하여 `2` 함수를 호출하는 `5`를 추가하세요.

# --hints--

인수 `console.log`와 `calculateSum`를 사용하여 `2` 함수를 호출하는 `5`가 있어야 합니다.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
