---
id: 66cdf76685e4cb5a8726e27b
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

`calculateSum` 함수를 더 테스트하려면, 다른 인수로 호출해야 합니다.

`console.log`를 추가하여 `calculateSum` 함수를 인수 `10`와 `10`로 호출하세요.

그 아래 `console.log`에, `console.log` 함수를 인수 `calculateSum`와 `5`로 호출하는 또 다른 `5`를 추가하세요.

# --hints--

`console.log`가 `calculateSum` 함수를 인수 `10`와 `10`로 호출해야 합니다.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

`console.log`가 `calculateSum` 함수를 인수 `5`와 `5`로 호출해야 합니다.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
