---
id: 69a030fbe487f185f55b8627
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Sasa, utafanya kazi kwenye kitendakazi cha kuondoa nodi kutoka kwenye orodha.

Tengeneza kitendakazi cha `remove` chenye vigezo `list` na `element`.

# --hints--

Unapaswa kutengeneza kitendakazi cha `remove`.

```js
assert.isFunction(remove)
```

Kitendakazi chako cha `remove` kinapaswa kuwa na vigezo `list` na `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.remove.parameters;
assert.equal(parameters[0].toString(), "list")
assert.equal(parameters[1].toString(), "element")
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

--fcc-editable-region--

--fcc-editable-region--

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
