---
id: 69a00cd4ec513345d40e95ec
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Jambo linalofuata utalofanya ni kitendakazi kinachokuwezesha kuangalia kama orodha ni tupu.

Ili kufanya hivyo, tengeneza kitendakazi kinachoitwa `isEmpty` kinachopokea kigezo `list`. Ndani ya kitendakazi hicho, rudisha ukaguzi wa kuona kama sifa ya `length` ya `list` ni sawa kabisa na `0`.

# --hints--

Unapaswa kutengeneza kitendakazi cha `isEmpty`.

```js
assert.isFunction(isEmpty)
```

Kitendakazi chako cha `isEmpty` kinapaswa kuwa na kigezo `list`.

```js
const regex = __helpers.functionRegex('isEmpty', ['list']);
assert.match(__helpers.removeJSComments(code), regex);
```

Kitendakazi chako cha `isEmpty` kinapaswa kurudisha `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
