---
id: 69a00a051b9a18694d00c16b
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, uta-tekeleza orodha iliyounganishwa.

Kuanza, tengeneza kitendakazi kinachoitwa `initList`. Ndani yake, rudisha kitu chenye sifa ya `head` iliyowekwa kuwa `null`, na sifa ya `length` yenye thamani ya `0`.

Huu ni kitendakazi utakachowito kuanzisha orodha iliyounganishwa. `head` inawakilisha kipengele cha kwanza katika orodha, mwanzoni, ni `null` kwa sababu hakuna kipengele katika orodha. `length` inafuatilia idadi ya nodi (vipengele) katika orodha, ambayo utaongeza au kupunguza unapoongeza au kuondoa nodi kutoka orodhani.

# --hints--

Unapaswa kutengeneza kitendakazi cha `initList`.

```js
assert.isFunction(initList)
```

Kitendakazi chako cha `initList` kinapaswa kurudisha kitu.

```js
assert.isObject(initList());
```

Kitendakazi chako cha `initList` kinapaswa kurudisha kitu chenye sifa ya `head` iliyowekwa kuwa `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Kitendakazi chako cha `initList` kinapaswa kurudisha kitu chenye sifa ya `length` iliyowekwa kuwa `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
