---
id: 69a04466917c66b132323a1b
title: Hatua 18
challengeType: 1
dashedName: step-18
---

# --description--

Kama `previous` si `null`, hii ina maana kipengele kinachotakiwa kuondolewa kiko si mwanzoni bali katikati ya orodha, kwa hivyo nodi ya sasa inapaswa kupitishwa kwa kuunganisha nodi iliyotangulia moja kwa moja na nodi inayofuata.

Ili kufanya hivyo, tengeneza tamko lingine la `if` linalokagua kama `previous` si `null` kabisa. Ikiwa ni hivyo, weka `previous.next` kuwa sawa na `current.next`.

# --hints--

Unapaswa kutengeneza tamko lingine la `if` lenye sharti la `previous !== null`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const removeSource = remove?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(removeSource);
const strippedRemove = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedRemove, 'if(previous!==null){');
```

Unapaswa kuweka `previous.next` kuwa `current.next`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const removeSource = remove?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(removeSource);
const strippedRemove = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedRemove, 'if(previous!==null){previous.next=current.next');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
