---
id: 69a04088e600bd8074a9852f
title: Hatua 17
challengeType: 1
dashedName: step-17
---

# --description--

Kama `current` ni `null`, ina maana hiyo ndio mwisho wa orodha, kwa hivyo kitendakazi kinapaswa kuacha kuendelea.

Tengeneza usemi wa `if` unaokagua kama `current` ni sawa kabisa na `null`. Ndani ya usemi wa `if`, tumia `return` kuondoka kwenye kitendakazi.

# --hints--

Unapaswa kutengeneza usemi wa `if` wenye sharti `current === null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*/);
```

Unapaswa kutumia neno kuu la `return` ndani ya usemi wako wa `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(current\s*===\s*null\s*\)\s*\{\s*return;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
