---
id: 69a00f673f171e29159e3ad8
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Nodi hii itahitaji sifa mbili, kipengele cha kuongeza, na kidokezo cha nodi inayofuata.

Ndani ya kitendakazi cha kuongeza, tengeneza kitu kinachoitwa `node`, kinapaswa kuwa na sifa inayoitwa `element`, na sifa ya `next` iliyowekwa kuwa `null`, ambayo ni kidokezo cha kipengele kinachofuata cha orodha.

# --hints--

Unapaswa kutengeneza kitu kinachoitwa `node` ndani ya kitendakazi chako cha `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Kitu chako cha `node` kinapaswa kuwa na sifa inayoitwa `element` iliyowekwa kuwa `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Kitu chako cha `node` kinapaswa kuwa na sifa inayoitwa `next` iliyowekwa kuwa `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
