---
id: 69a01864d343dc2f2484d131
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Baada ya kupata nodi ya mwisho, unahitaji kuunganisha nodi mpya nayo. Ili kufanya hivyo, weka `current.next` kuwa `node` chini ya mzunguko wa while lakini bado ndani ya kifungu cha `else`.

# --hints--

Unapaswa kuweka `current.next` kuwa `node` chini ya mzunguko wako wa `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
