---
id: 69a0102e761a245bb1cbd8e7
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Kwa sasa, unapaswa kujaribu orodha iliyounganishwa kwa uelewa bora.

Unda tofauti ya `myList` na uiweke kuwa `initList()`. Chini yake, andika wito wa `isEmpty()` ukitumia `myList` kama kigezo kwenye konsoli. Kisha, baada ya hapo, fanya wito wa `add` ukiwa na `myList` na nambari unayotaka, kisha andika `myList` kwenye konsoli. Chini ya hiyo, andika tena wito mwingine wa `isEmpty()` kwenye konsoli ukitumia `myList` kama kigezo tena.

# --hints--

Unapaswa kuunda tofauti ya `myList` ikiwa sawa na `initList()`.

```js
assert.deepEqual(myList, initList());
```

Unapaswa kuandika wito wa `isEmpty(myList)` kwenye konsoli.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(isEmpty(myList))');
```

Unapaswa kufanya wito wa `add` ukiwa na `myList` na nambari unayotaka kama vigezo.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedCode, /add\(myList,\d+\)/);
```

Unapaswa kuandika tofauti yako ya `myList` kwenye konsoli.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedCode, 'console.log(myList)');
```

Unapaswa kuandika tena `isEmpty(myList)` kwenye konsoli.

```js
const withoutComments = __helpers.removeJSComments(code);
const strippedCode = __helpers.removeWhiteSpace(withoutComments);
const regexMatch = strippedCode.matchAll(/console\.log\(isEmpty\(myList\)\)/g);
assert.isAtLeast([...regexMatch].length, 2);
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
