---
id: 69a03aaa1e6fe8152788ce6a
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Ndani ya kitendakazi cha `remove`, tumia `let` kuunda kigezo cha `previous` kilichowekwa kuwa `null`, na tena `let` kuunda kigezo cha `current` kilichowekwa kuwa `list.head`. Hivi ndizo vigezo utakavyotumia katika mzunguko utakaopita kwenye orodha.

`current` hufuata nodi inayokaguliwa kwa sasa katika orodha, kuanzia kichwani na kusonga mbele nodi moja kwa moja. `previous`, kwa upande mwingine, hufuata nodi iliyoko kabla ya ile ya sasa.

Hii ni muhimu wakati wa kuondoa nodi kwa sababu lazima usahihishe thamani inayofuata ya nodi ya awali ili iashirie nodi inayofuata ile inayotolewa, ili orodha iendelee kuunganishwa.

# --hints--

Unapaswa kutumia `let` kuunda kigezo cha `previous` kilichowekwa kuwa `null` ndani ya kitendakazi chako cha `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

Unapaswa kutumia `let` kuunda kigezo cha `current` kilichowekwa kuwa `list.head` ndani ya kitendakazi chako cha `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
