---
id: 69a014bd9b1730e94136279f
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Kama orodha si tupu, basi unataka kupata mwisho wake, ongeza kipengele kipya hapo, na uifanye iashirie `null`. Hii inamaanisha lazima upitie orodha kuanzia `head`.

Ili kuanza hilo, fungua kifungu cha `else`. Ndani yake, tumia `let` kuunda kigezo cha `current` na ukiseti kwa `list.head`.

# --hints--

Unapaswa kuwa na kifungu cha `else`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'else{');
```

Unapaswa kuunda kigezo cha `current` kilichosetwa kwa `list.head` ndani ya kifungu cha `else`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedAdd, /else\{(?:let|var)current=list\.head/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
