---
id: 69a013a86c97dac3784b2636
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa, kumbukumbu inapaswa kuonyesha kitu kama `{ head: { element: 42, next: null }, length: 1 }` kwenye konsoli.

Hata hivyo, hii siyo jinsi unavyotakavyoshughulikia kesi ambapo nodi iliyoongezwa ni kipengele cha kwanza kwenye orodha, yaani, kichwa.

Ili kufanya hivyo ipasavyo, tumia taarifa ya if kuangalia kama orodha ni tupu kwa kufanya wito wa kitendakazi `isEmpty` na `list`. Ikiwa hiyo ni kweli, basi hamisha mstari wako wa `list.head = node` ndani ya mwili wa taarifa ya `if`.

# --hints--

Unapaswa kuwa na taarifa ya `if` yenye sharti `isEmpty(list)`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*/);
```

Unapaswa kuhama mstari wako wa `list.head = node` ndani ya taarifa ya `if`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /if\s*\(isEmpty\(list\)\)\s*\{\s*list\.head\s*=\s*node;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
