---
id: 69a0179a85b00b6127e0af45
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Ndani ya mzunguko wa `while`, sahihisha `current` kuwa `current.next` ili mzunguko uendelee mbele kupitia orodha nodi moja kwa moja, na uishe wakati unafikia nodi yenye sifa ya `next` ya `null`, yaani, mwisho wa orodha.

# --hints--

Unapaswa kusahihisha thamani ya kigezo chako cha `current` kuwa `current.next` ndani ya mzunguko wa `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
