---
id: 69a012c6224fac85832247dd
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Umeona matokeo ya `true`, `{ head: null, length: 0 }`, na `true` kwenye konsoli? Hii hutokea kwa sababu hujawaweka thamani yoyote kwenye orodha ya `myList`, na hujapandisha thamani ya sifa ya `length`.

Endelea na weka thamani ya `node` kwa `list.head`, na chini ya hiyo, panda thamani ya sifa ya `length` ya `list`.

# --hints--

Unapaswa kuweka thamani ya `node` kwa `list.head`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

Unapaswa kupandisha thamani ya sifa ya `length` ya `list`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
