---
id: 69a045fd236d0d973a59f103
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Ikiwa `previous` ni `null`, hii ina maana kipengele kinachotakiwa kuondolewa kiko mwanzoni mwa orodha, kwa hivyo unapaswa kusahihisha `head` ili kuonyesha nodi inayofuata ili kuondoa kipengele cha kwanza.

Ili kufanya hivyo, ongeza kifungu cha `else`. Ndani yake, weka `list.head` kuwa `current.next`.

# --hints--

Unapaswa kuongeza kifungu cha `else` kwenye tamko lako la `if`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const removeSource = remove?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(removeSource);
const strippedRemove = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedRemove, '}else{');
```

Unapaswa kuweka `list.head` kuwa `current.next` ndani ya kifungu cha `else`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const removeSource = remove?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(removeSource);
const strippedRemove = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedRemove, '}else{list.head=current.next');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
