---
id: 69a00eaa43458e618331873c
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa utafanya kazi kwenye kitendakazi kitakachoongeza nodi kwenye orodha.

Tengeneza kitendakazi cha `add` chenye vigezo `list` na `element`.

# --hints--

Unapaswa kutengeneza kitendakazi cha `add`.

```js
assert.isFunction(add)
```

Kitendakazi chako cha `add` kinapaswa kuwa na vigezo `list` na `element`.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.add.parameters;
assert.equal(parameters[0].toString(), "list");
assert.equal(parameters[1].toString(), "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
