---
id: 69a016ac6edea7d8c8020c98
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa, tengeneza mzunguko wa `while` ukiwa na `current.next !== null` kama sharti. Hii inamaanisha mzunguko unapaswa kuendelea kuendeshwa mradi nodi ya sasa ina nodi nyingine baada yake.

# --hints--

Unapaswa kutengeneza mzunguko wa `while` ukiwa na sharti `current.next !== null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /while\s+\(current\.next\s*!==\s*null\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
