---
id: 69a0102e761a245bb1cbd8e7
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Kwa sasa, unapaswa kujaribu orodha iliyounganishwa kwa uelewa bora.

Unda tofauti ya `myList` na uiweke kuwa `initList()`. Chini yake, andika wito wa `isEmpty()` ukitumia `myList` kama kigezo kwenye konsoli. Kisha, baada ya hapo, fanya wito wa `add` ukiwa na `myList` na nambari unayotaka, kisha andika `myList` kwenye konsoli. Chini ya hiyo, andika tena wito mwingine wa `isEmpty()` kwenye konsoli ukitumia `myList` kama kigezo tena.

# --hints--

Unapaswa kuunda tofauti ya `myList` ikiwa sawa na `initList()`.

```js
assert.deepEqual(myList, initList());
```

Unapaswa kuandika wito wa `isEmpty(myList)` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Unapaswa kufanya wito wa `add` ukiwa na `myList` na nambari unayotaka kama vigezo.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Unapaswa kuandika tofauti yako ya `myList` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Unapaswa kuandika tena `isEmpty(myList)` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
