---
id: 66b102ee0bdbad34a9f85ba0
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви продовжите вивчати рядки, створюючи чатбота для вчителя.

Для початку додайте інструкцію `console` з повідомленням `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

У вас має бути інструкція `console`.

```js
assert.isAbove(logSpy.calls.length, 0);
```

Ваша інструкція для консолі має виводити повідомлення `"Hi there!"`.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
