---
id: 66b102ee0bdbad34a9f85ba0
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви продовжите вивчати рядки, створюючи чатбота для вчителя.

Для початку додайте інструкцію `console` з повідомленням `"Hi there!"`.

# --hints--

У вас має бути інструкція `console`.

```js
assert.match(code, /console\.log(.*)/);
```

Ваша інструкція для консолі має виводити повідомлення `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
