---
id: 66b59829dba144ff1351220f
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Тепер настав час створити привітання, використовуючи змінну `botName`.

У попередніх уроках ви навчились об’єднувати рядки за допомогою шаблонних літералів ось так:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Почніть зі створення змінної з назвою `greeting`.

Далі, використовуючи синтаксис шаблонних літералів, призначте рядок, який містить `My name is`, потім змінну `botName`, і завершується крапкою (`.`).

Нарешті, виведіть у консоль змінну `greeting`.

# --hints--

У вас має бути змінна з назвою `greeting`.

```js
assert.isNotNull(greeting);
```

Ви повинні призначити рядок змінній `greeting`.

```js
assert.isString(greeting);
```

Використайте шаблонні літерали, щоб об’єднати рядок `My name is` зі змінною `botName`, після чого поставте крапку (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

У вас має бути інструкція `console`, яка виводить у консоль змінну `greeting`.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
