---
id: 66b5928b86e907fdfbf34e56
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Тепер настав час призначити ім’я боту.

Створіть змінну з назвою `botName` і призначте їй рядкове значення `"teacherBot"`.

# --hints--

У вас має бути змінна з назвою `botName`.

```js
assert.isNotNull(botName);
```

Ваша змінна `botName` має бути рядком.

```js
assert.isString(botName);
```

Ваша змінна `botName` має мати значення `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
