---
id: 66b59b12a745e10011158f7b
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Наступний крок — створити ще кілька змінних, які будуть використані у майбутніх повідомленнях бота.

Створіть змінну з назвою `subject` і призначте їй рядкове значення `"JavaScript"`.

Потім створіть змінну з назвою `topic` і призначте їй рядкове значення `"strings"`.

# --hints--

У вас має бути змінна з назвою `subject`.

```js
assert.isNotNull(subject);
```

Ви повинні призначити рядок змінній `subject`.

```js
assert.isString(subject);
```

Ви повинні призначити рядок `"JavaScript"` змінній `subject`.

```js
assert.equal(subject, "JavaScript");
```

У вас має бути змінна з назвою `topic`.

```js
assert.isNotNull(topic);
```

Ви повинні призначити рядок змінній `topic`.

```js
assert.isString(topic);
```

Ви повинні призначити рядок `"strings"` змінній `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
