---
id: 66b59be6ab830800c4df9146
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер настав час використати змінні, які ви створили на попередньому кроці.

У попередніх уроках ви навчились працювати з шаблонними літералами так:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Почніть зі створення змінної `sentence`.

Використовуючи синтаксис шаблонних літералів, призначте рядок `Today, you will learn about [topic variable goes here] in [subject variable goes here].` змінній `sentence`.

Ви заміните заповнювачі `[topic variable goes here]` і `[subject variable goes here]` змінними `topic` і `subject` та синтаксисом `${}`.

Нарешті, виведіть змінну `sentence` у консоль.

# --hints--

У вас має бути змінна `sentence`.

```js
assert.isNotNull(sentence);
```

Ваша змінна `sentence` має бути рядком.

```js
assert.isString(sentence);
```

Ви маєте використати синтаксис шаблонних літералів, щоб призначити рядок `Today, you will learn about [topic variable goes here] in [subject variable goes here].` змінній `sentence`. Зверніться до прикладу, якщо потрібна допомога.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Ви маєте вивести змінну `sentence` у консоль.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
