---
id: 66b6f80fca500635d1e8af8d
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер настав час отримати доступ до останнього символу рядка.

Почніть із додавання ще однієї інструкції `console`, яка виводить повідомлення `Here is an example of accessing the last letter in the word [subject].`

Не забудьте замінити `[subject]` на фактичне значення змінної `subject` і використати правильний синтаксис шаблонних рядків.

# --hints--

Ваша інструкція `console` має виводити повідомлення `Here is an example of accessing the last letter in the word [subject].` Не забудьте використати правильний синтаксис шаблонних рядків.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+last\s+letter\s+in\s+the\s+word\s+\${subject}\.`\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

--fcc-editable-region--

--fcc-editable-region--
```
