---
id: 66b6fdb76441c738719039fa
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

У попередніх уроках ви навчились отримувати доступ до останнього символу рядка ось так:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` завжди дасть вам останній індекс рядка.

Створіть нову змінну з назвою `lastCharacter` і призначте їй значення останнього символу змінної `subject`.

Потім виведіть значення змінної `lastCharacter` у консоль.

# --hints--

У вас має бути змінна з назвою `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Ви маєте призначити значення останнього символу змінної `subject` змінній `lastCharacter`. Якщо потрібна допомога, зверніться до прикладу.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Ви маєте вивести значення змінної `lastCharacter` у `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
