---
id: 66b6d482bbb9e12f2e5ee1ae
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Для цієї наступної частини проєкту бот навчить працювати з властивістю рядка `length`.

Почніть зі створення нової змінної з назвою `strLengthIntro`.

Потім, використовуючи синтаксис шаблонних рядків, призначте рядок `Here is an example of using the length property on the word [subject].` змінній `strLengthIntro`.

Замініть `[subject]` на змінну `subject`, як ви робили раніше.

Нарешті, виведіть змінну `strLengthIntro` у консоль.

# --hints--

У вас має бути змінна з назвою `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

Ваша змінна `strLengthIntro` має бути рядком.

```js
assert.isString(strLengthIntro);
```

Ви маєте використати синтаксис шаблонних рядків, щоб призначити рядок `Here is an example of using the length property on the word [subject goes here].` змінній `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Ви маєте вивести змінну `strLengthIntro` у консоль.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
