---
id: 66b6e39031393e30f2c48d0f
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Щоб отримати довжину рядка, ви можете використати властивість <dfn>length</dfn>. Ця властивість повертає число символів у рядку.

Ось приклад:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Почніть із використання властивості `length`, щоб отримати довжину рядка `subject` і вивести це значення в консоль.

# --hints--


Ви повинні використати властивість `length`, щоб отримати довжину рядка `subject`.

```js
assert.match(code, /subject\.length/);
```

Ви повинні вивести довжину рядка `subject` у консоль.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
