---
id: 66b6ffb42b88e33943788abf
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Для останньої частини практичного заняття ви повторите, як знайти позицію індексу підрядка в рядку. Пам’ятайте, що підрядок — це частина рядка.

Почніть зі створення змінної з назвою `learningIsFunSentence` і призначте їй рядкове значення `"Learning is fun."`.

# --hints--

У вас має бути змінна з назвою `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

Ваш `learningIsFunSentence` має містити значення рядка.

```js
assert.isString(learningIsFunSentence);
```

Ви маєте призначити значення `"Learning is fun."` змінній `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
