---
id: 66b6efddeca35833cd6f0b03
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

У попередніх уроках ви навчились отримувати доступ до символів у рядку ось так:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Пам’ятайте, що індекси починаються з `0`, тому перша літера в рядку завжди буде за індексом `0`.

Почніть із додавання ще однієї інструкції `console`.

Всередині інструкції `console` виведіть першу літеру змінної `subject` за допомогою дужкової нотації та правильного індексу.

# --hints--

Ваша інструкція `console` має вивести першу літеру змінної `subject` за допомогою дужкової нотації та індексу `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
