---
id: 66b6e62423e8a031d6c1f03d
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Тепер настав час отримати довжину рядка `topic`.

Ви можете використовувати шаблонні літерали всередині інструкцій `console` ось так:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Почніть з виведення повідомлення `Here is an example of using the length property on the word [topic].` у консоль.

Не забудьте замінити `[topic]` на змінну `topic` і використати правильний синтаксис шаблонних літералів, як ви робили на попередніх кроках.

Потім додайте другу інструкцію `console.log`, яка виводить довжину рядка `topic` у консоль.

# --hints--

Ви повинні використати синтаксис шаблонних літералів, щоб вивести повідомлення `Here is an example of using the length property on the word [topic].` у `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

У вас має бути друга інструкція `console`, щоб вивести довжину змінної `topic` у `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
