---
id: 66b6f586767a1534f3097353
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Тепер настав час отримати доступ до другої літери змінної `subject`.

Почніть із додавання інструкції `console`, яка виводить повідомлення `Here is an example of accessing the second letter in the word [subject].`. Не забудьте замінити `[subject]` на фактичне значення змінної `subject` і використати правильний синтаксис шаблонного рядка.

Потім додайте ще одну інструкцію `console`, яка виводить другу літеру змінної `subject` за допомогою дужкової нотації та правильного індексу.

# --hints--

Ваша інструкція `console` має виводити повідомлення `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

У вас має бути ще одна інструкція `console`, яка виводить другу літеру змінної `subject` за допомогою дужкової нотації з індексом `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
