---
id: a3bfc1673c0526e06d3ac698
title: بناء حاسبة مجموع الأعداد الأولية
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

في هذا المختبر، ستبني حاسبة تأخذ رقمًا وتُرجع مجموع كل الأعداد الأولية التي تكون أقل من أو تساوي ذلك الرقم.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة `sumPrimes` تقبل رقمًا كمعلمة.
1. يجب أن تُرجع دالة `sumPrimes` مجموع كل الأعداد الأولية التي تكون أقل من أو تساوي الرقم المُدخل.
1. إذا كان الرقم المُدخل أقل من `2`، يجب أن تُرجع الدالة `0`.

# --hints--

يجب أن تملك دالة `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

يجب أن تُرجع `sumPrimes(10)` القيمة `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

يجب أن تُرجع `sumPrimes(5)` القيمة `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

يجب أن تُرجع `sumPrimes(2)` القيمة `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

يجب أن تُرجع `sumPrimes(0)` القيمة `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

يجب أن تُرجع `sumPrimes(977)` القيمة `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
