---
id: 646c48df8674cf2b91020ecc
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Se te ha proporcionado un modelo básico. Comienza agregando un elemento `link` dentro del elemento `head`. Para ese elemento `link`, establece el atributo `rel` en `stylesheet` y el `href` en `./styles.css`.

# --hints--

Debes tener un elemento `link`.

```js
assert.exists(document.querySelector("head>link"));
```

El elemento `link` debe tener un atributo `rel` y un atributo `href` establecido en `stylesheet` y `./styles.css`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
console.log(href)
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
</head>
--fcc-editable-region--
<body></body>
</html>
```

```css

```
