---
id: 6476f5c17f99146071ee884c
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Podría ver que no pasa nada. El elemento `.cat-head` no se movió a pesar de establecer `top` y `left` en `100px` cada uno. Pero eso no es el caso con el posicionamiento `relative`.

Cuando usa el valor `relative`, el elemento todavía se posiciona de acuerdo con el flujo normal del documento, pero los valores de `top`, `left`, `bottom` y `right` se activan.

En lugar de `static`, da a tu `.cat-head` una posición de `relative`, y deja tanto las propiedades `top` como `left` tal como están.

# --hints--

Tu selector `.cat-head` debe tener una propiedad `position` establecida en `relative`. Asegúrate de añadir un punto y coma.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position, 'relative');
```

Tu selector `.cat-head` debe tener una propiedad `top` establecida en `100px`. Asegúrate de añadir un punto y coma.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top, '100px');
```

Tu selector `.cat-head` debe tener una propiedad `left` establecida en `100px`. Asegúrate de añadir un punto y coma.

```js
assert.equal(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left, '100px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
