---
id: 66a37f37ef5823a313de8c26
title: Construye una Órbita Lunar
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

En este laboratorio, crearás una animación simple de la órbita de la Luna alrededor de la Tierra usando HTML y CSS. La Tierra estará en el centro del sistema, y la Luna girará a su alrededor.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Historia de usuario:**

1. Debes alinear todos los elementos al centro de la página estableciendo la altura al `100%` del viewport y usando un diseño flexbox en el elemento `body`.

2. Debes tener un `div` con la clase `space`.

3. Dentro del elemento `.space`, debe haber dos elementos `div` más con la clase `earth` y `orbit`, respectivamente, en ese orden.

4. Dentro del elemento `.orbit`, debe haber otro `div` con la clase `moon`.

5. El elemento `div` con una `class` de `space` debe estar centrado en la página y tener un ancho y alto de `200px`.

6. El elemento `div` con una `class` de `space` debe usar posicionamiento `relative`.

7. El elemento `.earth` debe usar posicionamiento `absolute`. Posiciona el centro del elemento `.earth` en el punto medio de su padre, tanto en los ejes vertical (superior) como horizontal (izquierdo). Después de eso, desplaza el elemento `.earth` hacia atrás en la mitad de su ancho y alto, para centrarlo dentro de su padre, el elemento `.space`.

8. El elemento `.earth` debe tener un ancho y alto de `100px`.

9.  El elemento `.orbit` debe tener un ancho y alto de `200px`.

10.  El elemento `.orbit` debe ser posicionado usando posicionamiento `absolute`. Su esquina inferior derecha debe estar en el centro del elemento `.space` usando una propiedad `transform` que lo desplace un `-50%` en ambos ejes vertical y horizontal.

11. La trayectoria de la órbita para la luna alrededor de la Tierra debe ser un círculo.

12. El elemento `.moon` debe ser posicionado usando posicionamiento `absolute` y tener un ancho y alto de `30px`. El elemento `.moon` debe posicionarse en la parte superior del elemento `.orbit` y estar centrado horizontalmente.

13. Debes ajustar aún más la posición horizontal del elemento `.moon` moviendo el elemento hacia la izquierda por la mitad de su ancho.

14. The `.earth` selector should have a background color and a `border-radius` of `50%`.

15. The `.moon` selector should have a background color and a `border-radius` of `50%`.

16. Debes definir una animación `@keyframes orbit` que rote el elemento `.orbit` 360 grados alrededor de su centro. Debes aplicar esta animación al elemento `.orbit` con una duración de `5` segundos, una función de temporización lineal y iteraciones infinitas.

**Nota:** Asegúrate de vincular tu hoja de estilos en tu HTML y aplicar tu CSS.

# --hints--

El contenido del elemento `body` debe estar centrado en la página estableciendo la altura a `100%` de la altura del viewport, `display` a `flex`, `justify-content` y `align-items` a `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Debes tener un `div` principal con la clase `space`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

El elemento `.space` debe tener un ancho y alto de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

El elemento `.space` debe tener una propiedad `position` de `relative`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Debes tener un `div` con una clase de `earth` dentro del `div` con una `class` de `space`.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

El elemento `.earth` debe ser el primer hijo del elemento `.space`.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

El elemento `.earth` debe tener un ancho y alto de `100px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

El elemento `.earth` debe usar posicionamiento `absolute`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

El elemento `.earth` debe tener una posición `top` y `left` de `50%`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

El elemento `.earth` debe moverse hacia atrás y arriba `50%` configurando la propiedad `transform` a `translate(-50%, -50%)`.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Debes tener un `div` con la clase `orbit` dentro del elemento `.space`.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

El elemento `.orbit` debe venir después del elemento `.earth`.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

El elemento `.orbit` debe tener un ancho y alto de `200px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

El elemento `.orbit` debe usar posicionamiento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

El elemento `.orbit` debe estar posicionado en el centro del elemento `.space` usando una propiedad `transform` que lo desplace `-50%` en ambos ejes vertical y horizontal.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Debes tener un `div` con la clase `moon` dentro del elemento `.orbit`.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

El elemento `.moon` debe tener un ancho y alto de `30px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

El elemento `.moon` debe ser posicionado usando posicionamiento `absolute`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

El borde superior del elemento `.moon` debe estar alineado con el borde superior de su bloque contenedor.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Debes mover el borde izquierdo del elemento `.moon` al centro del bloque contenedor (`50%` desde el lado izquierdo).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

El elemento `.moon` debe ser ajustado horizontalmente al traducirlo a `-50%` a lo largo del eje x.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.transform, 'translateX(-50%)');
```

The `.earth` selector should have a background color.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

The `.earth` selector should have `border-radius` of `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

The `.moon` selector should have a background color.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

The `.moon` selector should have `border-radius` of `50%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Debes tener una regla `@keyframes`.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

Tu nueva regla `@keyframes` debe llamarse `orbit`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

Tu regla `@keyframes orbit` debe tener un selector `0%`.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

Tu regla `@keyframes orbit` debe tener un selector `100%`.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Tu selector `0%` debe tener una propiedad `transform` configurada como `rotate(0deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Tu selector `100%` debe encontrarse después que tu selector `0%`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Tu selector `100%` debe tener una propiedad `transform` configurada como `rotate(360deg) translate(-50%, -50%)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Tu selector `.orbit` debe tener una propiedad `animation` configurada como `orbit 5s linear infinite`.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

