---
id: 6989f3632f8b8bdf168aa321
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Ndani ya mzunguko wako wa `while`, tumia `queue.shift()` kuondoa na kupata kipengele cha kwanza kutoka kwenye foleni. Hii inatekeleza tabia ya kwanza-kuingia-kwa-kwanza-kutoka (FIFO) inayojulikana katika BFS.

Gawanya safu ya data hii katika vigezo vitatu: `current`, `opensUsed`, na `closesUsed`.

# --hints--

Unapaswa kutumia `queue.shift()` kupata kipengele cha kwanza kutoka kwenye foleni.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

Unapaswa kugawanya matokeo hayo katika vigezo vitatu: `current`, `opensUsed`, na `closesUsed`.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}
```
