---
id: 6989f207306332c06b0add39
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Kwa njia ya utafutaji wa upana kwanza, utatumia foleni kufuatilia hali tofauti unapounda mchanganyiko wa mabano ya kawaida. Kila hali itawakilishwa kama safu ya data yenye vipengele vitatu:

- Mfuatano wa herufi unaoendelea kujengwa
- Idadi ya mabano ya kufungua yaliyotumika hadi sasa
- Idadi ya mabano ya kufunga yaliyotumika hadi sasa

Tengeneza kigezo kinachoitwa `queue` na uanze na safu ya data yenye safu moja: `['', 0, 0]`. Hii inawakilisha hali ya kuanzia yenye mfuatano wa herufi tupu na mabano sifuri yaliyotumika.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `queue` katika kitendakazi chako cha `genParentheses`.

```js
assert.match(genParentheses.toString(), /queue/);
```

Unapaswa kuanzisha `queue` kwa safu ya data yenye safu `['', 0, 0]`.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
