---
id: 6989f2e35e6941e13e445156
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Sasa utaekeleza mzunguko mkuu wa BFS. Unda mzunguko wa `while` unaoendelea mradi `queue` haijajaa, yaani, `queue.length` ni kubwa kuliko `0`.

Ndani ya mzunguko huo, andika `queue` kwenye konsoli.

# --hints--

Unapaswa kuunda mzunguko wa `while` wenye sharti la `queue.length > 0`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

Unapaswa kuandika `queue` kwenye konsoli ndani ya mzunguko wako wa `while`.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
