---
id: 6989eb65324d97775c4965fe
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, uta-tekeleza kitendakazi kinachozalisha mchanganyiko wote halali wa mabano ya kawaida kwa kutumia mbinu ya utafutaji wa upana kwanza (BFS). Kwa mfano, mchanganyiko halali wa mabano ya kawaida ya jozi mbili ni `(())` na `()()`.

Anza kwa kuunda kitendakazi kinachoitwa `genParentheses` chenye kigezo kimoja `pairs`. Kwa sasa, rudisha safu ya data tupu kutoka kwa kitendakazi hicho.

# --hints--

Unapaswa kufafanua kitendakazi kinachoitwa `genParentheses`.

```js
assert.isFunction(genParentheses);
```

Kitendakazi chako cha `genParentheses` kinapaswa kuwa na kigezo kimoja kinachoitwa `pairs`.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

Kitendakazi chako kinapaswa kurudisha safu ya data tupu.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
