---
id: 6989f207306332c06b0add39
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Kwa njia ya utafutaji wa upana kwanza, utatumia foleni kufuatilia hali tofauti unapounda mchanganyiko wa mabano ya kawaida. Kila hali itawakilishwa kama safu ya data yenye vipengele vitatu:

- Mfuatano wa herufi unaojengwa kwa sasa
- Idadi ya mabano ya kawaida ya kufungua yaliyotumika hadi sasa
- Idadi ya mabano ya kawaida ya kufunga yaliyotumika hadi sasa

Tengeneza kigezo kinachoitwa `queue` na uanze na safu ya data yenye safu moja: `['', 0, 0]`. Hii inawakilisha hali ya kuanzia yenye mfuatano wa herufi tupu na mabano sifuri yaliyotumika.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `queue` katika kitendakazi chako cha `genParentheses`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
assert.exists(genParentheses?.variables.queue);
```

Unapaswa kuanzisha `queue` kwa safu ya data yenye safu `['', 0, 0]`.

```js
const explorer = await __helpers.Explorer(code);
const { genParentheses } = explorer.allFunctions;
const { queue } = genParentheses?.variables ?? {};
assert.isFalse(queue?.value.isEmpty());
const queueValue = eval(queue?.value.toString());
assert.deepEqual(queueValue, [['', 0, 0]]);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  
  return result;
}
```
