---
id: 6989f48c3f0292c84f8d8938
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa unahitaji kubaini wakati umejenga mchanganyiko kamili wa mabano ya kawaida. Mchanganyiko kamili una urefu sawa na mara mbili ya nambari ya jozi (kwa sababu kila jozi huchangia mabano ya kufungua na kufunga).

Ndani ya mzunguko wako wa `while`, ongeza tamko la `if` ili ukague kama `current.length === 2 * pairs`. Wakati hali hii ni kweli, sukuma `current` kwenye safu ya data ya `result`.

# --hints--

Unapaswa kuwa na tamko la `if` lenye hali ya `current.length === 2 * pairs` ndani ya mzunguko wako wa `while`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

Unapaswa kusukuma `current` kwenye safu ya data ya `result` ndani ya tamko lako jipya la `if`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
  let queue = [['', 0, 0]];
  let result = [];
  
  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
