---
id: 6989ef885db3a3488a13b426
title: Hatua ya 2
challengeType: 1
dashedName: step-2
---

# --description--

Kabla ya kutekeleza algoriti kuu, unahitaji kuthibitisha ingizo. Kigezo cha `pairs` kinapaswa kuwa nambari kamili, kwani kinawakilisha idadi ya jozi za mabano ya kawaida za kuzalisha.

Ongeza taarifa ya `if` mwanzoni mwa kitendakazi chako ili kukagua kama `pairs` si nambari kamili. Tumia njia ya `Number.isInteger()` kwa ajili ya hilo.

Kama hali ni kweli, rudisha mfuatano wa herufi `The number of pairs should be an integer`.

# --hints--

Unapaswa kuwa na taarifa ya if yenye hali `!Number.isInteger(pairs)` ndani ya kitendakazi cha `genParentheses`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

Unapaswa kurudisha mfuatano wa herufi wakati ingizo si nambari kamili.

```js
let result = genParentheses("i");
assert.isString(result);
```

Unapaswa kurudisha mfuatano wa herufi `The number of pairs should be an integer` wakati ingizo ni mfuatano wa herufi.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```

