---
id: 6989f1ba358acba651d5ee6b
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sasa utaanzisha muundo wa data kuhifadhi matokeo yako. Tengeneza kigezo kinachoitwa `result` na uanze na safu ya data tupu. Safu hii ya data itahifadhi mchanganyiko wote halali wa mabano ya kawaida utakayozalisha.

Sahihisha taarifa yako ya return ili irudishe `result` badala ya safu ya data tupu.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `result` katika kitendakazi chako cha `genParentheses`.

```js
assert.match(genParentheses.toString(), /result/);
```

Unapaswa kuanzisha `result` kuwa safu ya data tupu.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

Unapaswa kurudisha `result` mwishoni mwa kitendakazi chako.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
