---
id: 6989f10a72ad424d4ad6d514
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Ifuatayo, unahitaji kuthibitisha kuwa idadi ya jozi ni angalau moja, kwa sababu huwezi kuunda mchanganyiko wa mabano ya kawaida ukiwa na jozi sifuri au hasi.

Ongeza tamko lingine la `if` ili ukague kama `pairs` ni ndogo kuliko `1`. Ikiwa hali hii ni kweli, rudisha mfuatano wa herufi `The number of pairs should be at least 1`.

# --hints--

Unapaswa kuwa na tamko la pili la `if` lenye hali `pairs < 1`.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

Unapaswa kurudisha mfuatano wa herufi `The number of pairs should be at least 1` wakati `pairs` ni ndogo kuliko `1`.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  
  return [];
}
```
