---
id: 696920c0c98a1ed58eb86293
title: 构建产品展示
challengeType: 25
dashedName: lab-product-showcase
demoType: onClick
---

# --description--

在本实验中，你已经获得了所有的超文本标记语言和 CSS。你将使用 TypeScript 完成一个产品展示。

**目标：** 完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该定义一个名为 `Item` 的接口，包含：
   - 一个 `type` 属性，该属性只能是 `"book"`、`"electronics"` 或 `"clothing"`。
   - 一个类型为 `id` 的 `string` 属性。
   - 一个类型为 `price` 的 `number` 属性。

1. 你应该定义一个名为 `Book` 的接口，继承自 `Item`。它的 `type` 属性应设置为字面值 `"book"`。它还应有类型为 `title` 的 `author` 和 `string` 属性。

1. 你应该定义一个名为 `Electronics` 的接口，继承自 `Item`。它的 `type` 属性应设置为字面值 `"electronics"`。它还应有类型为 `item` 的 `model` 和 `string` 属性，以及一个可选的类型为 `warranty` 的 `number` 属性。

1. 你应该定义一个名为 `Clothing` 的接口，继承自 `Item`。它的 `type` 属性应设置为字面值 `"clothing"`。它还应有类型为 `item` 的 `brand` 和 `string` 属性，以及一个可选的 `size` 属性，该属性只能是 `"S"`、`"M"` 或 `"L"`。

1. 你应该定义一个名为 `Product` 的并集类型，组合 `Book`、`Electronics` 和 `Clothing` 接口。

1. 你应该定义一个通用的类，名为 `Collection`，带有类型参数 `T`。构造函数应接受一个类型为 `T` 的数组，并将其保存到名为 `items` 的属性中。

1. `Collection` 类应有一个 `getAll()` 方法，返回集合中的所有项。

1. `Collection` 类应有一个 `filter()` 方法，接受一个回调函数，并返回 `items` 数组中满足给定条件的元素。回调函数应接收类型为 `T` 的元素并返回布尔值。

1. 你应该定义一个名为 `renderProduct` 的函数，带有类型为 `Product` 的参数。它应返回一个表示产品的超文本标记语言字串，包含：
   - 一个类名为 `item` 且 ID 为产品 `id` 值的元素。
   - 一个类名为 `price` 且文本为产品 `price` 的元素。
   - 一个包含额外产品信息的元素。

1. `renderProduct` 函数应使用类型缩小，根据产品类型渲染不同的超文本标记语言：

   - 对于书籍，渲染的超文本标记语言应符合 `Book: [title] by [author]` 格式，其中 `[title]` 和 `[author]` 分别替换为标题和作者。
   - 对于电子产品，渲染的超文本标记语言应符合 `Electronics: [item] - [model]` 格式。如果产品有保修，应追加 ` - Warranty: [warranty] year(s)`。将 `[item]`、`[model]` 和 `[warranty]` 替换为项、型号和保修。
   - 对于服装，渲染的超文本标记语言应符合 `Clothing: [item] by [brand]` 格式。如果产品有尺码，应追加 ` - Size [size]`。将 `[item]`、`[brand]` 和 `[size]` 替换为项、品牌和尺码。
   - 如果遇到无效的产品类型，应抛出格式为 `Unknown product type: [product]` 的误差，其中 `[product]` 替换为无效的产品对象。使用 `JSON.stringify` 将无效产品转换为字串。

1. 你应该创建一个 `Collection<Product>` 实例，至少包含每种产品类型的一项，并将该实例赋值给变量 `products`。

1. 你应该定义一个 `showProducts` 函数，在 `#output` 元素中显示产品。它应接受一个可选的过滤器，表示产品类型；当省略时显示所有产品；当提供时只显示该类型的产品。

1. 你应通过调用 `renderProduct` 函数将每个产品转换为超文本标记语言字串，并将合并结果注入 `#output` 元素来渲染产品。

1. 点击 `#all`、`#books`、`#electronics` 和 `#clothing` 按钮应在 `#output` 元素中显示对应的产品集合。

1. 页面加载时，应默认显示所有产品。

# --hints--

你应该有一个名为 `Item` 的接口。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Item);
```

你的 `Item` 接口应有一个 `type` 属性，该属性只能是 `"book"`、`"electronics"` 或 `"clothing"`。

```js
const explorer = await __helpers.Explorer(code);
const { Item } = explorer.interfaces;
const typeProp = Item.typeProps.type;
assert.isTrue(typeProp.isUnionOf(["'book'", "'electronics'", "'clothing'"]));
```

你的 `Item` 接口应有一个类型为 `id` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "id", type: "string"}));
```

你的 `Item` 接口应有一个类型为 `price` 的 `number` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Item.hasTypeProps({name: "price", type: "number"}));
```

你应该有一个名为 `Book` 的接口。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Book);
```

你的 `Book` 接口应继承自 `Item`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.doesExtend("Item"));
```

你的 `Book` 接口应有一个 `type` 属性，值为 `"book"`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "type", type: "'book'"}));
```

你的 `Book` 接口应有一个类型为 `title` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "title", type: "string"}));
```

你的 `Book` 接口应有一个类型为 `author` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Book.hasTypeProps({name: "author", type: "string"}));
```

你应该有一个名为 `Electronics` 的接口。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Electronics);
```

你的 `Electronics` 接口应继承自 `Item`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.doesExtend("Item"));
```

你的 `Electronics` 接口应有一个 `type` 属性，值为 `"electronics"`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "type", type: "'electronics'"}));
```

你的 `Electronics` 接口应有一个类型为 `item` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "item", type: "string"}));
```

你的 `Electronics` 接口应有一个类型为 `model` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "model", type: "string"}));
```

你的 `Electronics` 接口应有一个可选的类型为 `warranty` 的 `number` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Electronics.hasTypeProps({name: "warranty", type: "number", isOptional: true}));
```

你应该有一个名为 `Clothing` 的接口。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.interfaces.Clothing);
```

你的 `Clothing` 接口应继承自 `Item`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.doesExtend("Item"));
```

你的 `Clothing` 接口应有一个 `type` 属性，值为 `"clothing"`。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "type", type: "'clothing'"}));
```

你的 `Clothing` 接口应有一个类型为 `item` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "item", type: "string"}));
```

你的 `Clothing` 接口应有一个类型为 `brand` 的 `string` 属性。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.interfaces.Clothing.hasTypeProps({name: "brand", type: "string"}));
```

你的 `Clothing` 接口应有一个可选的 `size` 属性，该属性只能是 `"S"`、`"M"` 或 `"L"`。

```js
const explorer = await __helpers.Explorer(code);
const { Clothing } = explorer.interfaces;
assert.isTrue(Clothing.hasTypeProps({name: "size", isOptional: true}));
assert.isTrue(Clothing.typeProps.size.isUnionOf(["'S'", "'M'", "'L'"]));
```

你应该定义一个名为 `Product` 的并集类型，组合 `Book`、`Electronics` 和 `Clothing` 接口。

```js
const explorer = await __helpers.Explorer(code);
const { Product } = explorer.types;
assert.isTrue(Product.isUnionOf(["Book", "Electronics", "Clothing"]));
```

你应该有一个名为 `Collection` 的类，带有类型参数 `T`。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.isTrue(Collection.typeParameters[0].matches("T"));
```

你的 `Collection` 类的构造函数应接受一个类型为 `T` 的数组，并将其保存到名为 `items` 的属性中。

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.items, arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.items, arr);
```

你的 `Collection` 类应有一个 `getAll` 方法。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
console.log(Collection.methods)
console.log(Collection.methods.getAll)
assert.exists(Collection.methods.getAll);
```

你的 `getAll` 方法应返回集合中的所有项。

```js
let arr = [1, 2, 3];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.getAll(), arr);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.getAll(), arr);
```

你的 `Collection` 类应有一个 `filter` 方法。

```js
const explorer = await __helpers.Explorer(code);
const { Collection } = explorer.classes;
assert.exists(Collection.methods.filter);
```

你的 `filter` 方法应接受一个回调函数，并返回一个数组，包含 `items` 属性中满足给定条件的元素。

```js
let arr = [1, 2, 3, 4, 5, 6];
const numCollection = new Collection(arr);
assert.deepEqual(numCollection.filter(n => n % 2 === 0), [2, 4, 6]);
arr = ["a", "b", "c", "d"];
const stringCollection = new Collection(arr);
assert.deepEqual(stringCollection.filter(s => s === "c"), ["c"]);
```

你应该有一个名为 `renderProduct` 的函数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.renderProduct);
```

`renderProduct(product)` 应返回一个字串，包含一个类名为 `item` 且 ID 等于 `product.id` 的元素。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  assert.isNotNull(testEl.querySelector(`[class="item"][id="${p.id}"]`));
})
```

`renderProduct(product)` 应返回一个字串，包含一个类名为 `price` 且文本等于 `product.price` 的元素。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const price = testEl.querySelector(`[class="price"]`);
  assert.include(price.innerText, p.price.toString());
})
```

当 `renderProduct` 被调用，传入 `Book` 类型的产品 `p` 时，应返回一个字串，包含一个文本格式为 `Book: [title] by [author]` 的元素。将 `[title]` 和 `[author]` 替换为 `p` 的对应属性。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },
  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const expectedText = `Book: ${p.title} by ${p.author}`;
  assert.include(testEl.innerText, expectedText);
})
```

当 `renderProduct` 被调用，传入 `Electronics` 类型的产品 `p` 时，应返回一个字串，包含一个文本格式为 `Electronics: [item] - [model]` 的元素；如果 `Electronics: [item] - [model] - Warranty: [warranty] year(s)` 不为 undefined，则格式为 `p.warranty`。将 `[item]`、`[model]` 和 `[warranty]` 替换为 `p` 的对应属性。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Electronics: ${p.item} - ${p.model}`;
  const expectedText = !p.warranty ? baseFormat : baseFormat + ` - Warranty: ${p.warranty} year(s)`;
  assert.include(testEl.innerText, expectedText);
  if (!p.warranty) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

当 `renderProduct` 被调用，传入 `Clothing` 类型的产品 `p` 时，应返回一个字串，包含一个文本格式为 `Clothing: [item] by [brand]` 的元素；如果 `Clothing: [item] by [brand] - Size [size]` 不为 undefined，则格式为 `p.size`。将 `[item]`、`[brand]` 和 `[size]` 替换为 `p` 的对应属性。

```js
const testEl = document.createElement("div");
const productObjects = [
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 }

]
productObjects.forEach(p => {
  const renderedProduct = renderProduct(p);
  testEl.innerHTML = renderedProduct;
  const baseFormat = `Clothing: ${p.item} by ${p.brand}`;
  const expectedText = !p.size ? baseFormat : baseFormat + ` - Size ${p.size}`;
  assert.include(testEl.innerText, expectedText);
  if (!p.size) {
    assert.notInclude(testEl.innerText, `undefined`);
  }
})
```

当 `renderProduct` 被调用，传入无效产品时，应抛出格式为 `Unknown product type: [p]` 的误差。将 `[p]` 替换为使用 `JSON.stringify` 字符串化的无效产品。

```js
const invalidProduct = { item: "T-Shirt", price: 19.99 };
const expected = `Unknown product type: ${JSON.stringify(invalidProduct)}`;
try {
  renderProduct(invalidProduct);
  assert.fail();
} catch({name, message}) {
  if (name == "AssertionError") {
    assert.fail("renderProduct should throw with invalid product");
  } else {
    assert.equal(message, expected);
  }
}
```

你应该创建一个 `Collection<Product>` 实例，并赋值给变量 `products`。

```js
assert.instanceOf(products, Collection);
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.products.value.typeArguments[0].matches("Product"));
```

你的 `products` 变量应至少包含每种产品类型的一项。

```js
assert.isTrue(products.items.some(p => p.type === "book"));
assert.isTrue(products.items.some(p => p.type === "electronics"));
assert.isTrue(products.items.some(p => p.type === "clothing"));
```

你应该有一个名为 `showProducts` 的函数。

```js
assert.isFunction(showProducts);
```

点击 `#all` 按钮应在 `#output` 元素中显示所有产品。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
books.dispatchEvent(new Event("click"));
all.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

点击 `#books` 按钮应在 `#output` 元素中显示对应的产品集合。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "book") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
books.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

点击 `#electronics` 按钮应在 `#output` 元素中显示对应的产品集合。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "electronics") {
    count++;
  }
});
clothing.dispatchEvent(new Event("click"));
electronics.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

点击 `#clothing` 按钮应在 `#output` 元素中显示对应的产品集合。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
let count = 0;
products.items.forEach(i => {
  if (i.type === "clothing") {
    count++;
  }
});
books.dispatchEvent(new Event("click"));
clothing.dispatchEvent(new Event("click"));
assert.lengthOf(output.querySelectorAll(".item"), count);
```

页面加载时应默认显示所有产品。使用 `DOMContentLoaded` 事件确保这一点。

```js
const len = products.items.length;
assert.isAtLeast(len, 3);
document.dispatchEvent(new Event("DOMContentLoaded"));
assert.lengthOf(output.querySelectorAll(".item"), len);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  <title>Product Showcase</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Product Showcase</h1>
  <div class="buttons">
    <button id="all">All</button>
    <button id="books">Books</button>
    <button id="electronics">Electronics</button>
    <button id="clothing">Clothing</button>
  </div>
  <div id="output" class="product-list"></div>
  <script src="index.ts"></script>
</body>
</html>
```

```css
body {
  font-family: 'Segoe UI', Tahoma, Geneva, Verdana, sans-serif;
  background: linear-gradient(135deg, #0f4c75 0%, #3282b8 100%);
  padding: 40px 20px;
  margin: 0;
  color: #1a1a1a;
  min-height: 100vh;
}

h1 {
  text-align: center;
  margin-bottom: 40px;
  color: #ffffff;
  font-size: 2.5rem;
  font-weight: 600;
  letter-spacing: -0.5px;
}

.buttons {
  display: flex;
  justify-content: center;
  gap: 12px;
  margin-bottom: 40px;
  flex-wrap: wrap;
}

button {
  padding: 12px 24px;
  font-size: 15px;
  font-weight: 500;
  cursor: pointer;
  border: none;
  border-radius: 8px;
  background: rgba(255, 255, 255, 0.9);
  color: #333;
  transition: all 0.3s ease-in-out;
  box-shadow: 0 2px 8px rgba(0, 0, 0, 0.1);
}

button:hover {
  background: white;
  transform: translateY(-2px);
  box-shadow: 0 4px 12px rgba(0, 0, 0, 0.15);
}

button.active {
  background: #ff9800;
  color: #1a1a1a;
  box-shadow: 0 4px 15px rgba(255, 152, 0, 0.4);
  font-weight: 600;
}

#output {
  display: grid;
  grid-template-columns: repeat(auto-fill, minmax(300px, 1fr));
  gap: 24px;
  max-width: 1200px;
  margin: 0 auto;
}

.item {
  background: white;
  border-radius: 12px;
  padding: 24px;
  box-shadow: 0 4px 15px rgba(0, 0, 0, 0.1);
  transition: all 0.3s ease-in-out;
  display: flex;
  flex-direction: column;
  gap: 12px;
}

.item:hover {
  transform: translateY(-8px);
  box-shadow: 0 12px 25px rgba(0, 0, 0, 0.15);
}

.item strong {
  font-size: 12px;
  text-transform: uppercase;
  letter-spacing: 1px;
  color: #0f4c75;
  font-weight: 700;
}

.item > div:not(.price) {
  font-size: 16px;
  color: #555;
  line-height: 1.5;
}

.price {
  margin-top: auto;
  font-size: 24px;
  font-weight: 700;
  color: #0f4c75;
  padding-top: 12px;
  border-top: 2px solid #f0f0f0;
}

@media (max-width: 768px) {
  #output {
    grid-template-columns: repeat(auto-fill, minmax(250px, 1fr));
    gap: 16px;
  }

  h1 {
    font-size: 2rem;
  }

  .item {
    padding: 16px;
  }
}

@media (max-width: 480px) {
  #output {
    grid-template-columns: 1fr;
  }

  button {
    padding: 10px 16px;
    font-size: 14px;
  }

  .price {
    font-size: 20px;
  }
}
```

```ts
interface Item {
  type: "book" | "electronics" | "clothing";
  id: string;
  price: number;
}

interface Book extends Item {
  type: "book";
  title: string;
  author: string;
}

interface Electronics extends Item {
  type: "electronics";
  item: string;
  model: string;
  warranty?: number;
}

interface Clothing extends Item {
  type: "clothing";
  item: string;
  brand: string;
  size?: "S" | "M" | "L";
}

type Product = Book | Electronics | Clothing;

class Collection<T> {
  items: T[];
  constructor(items: T[]) {
    this.items = items;
  }

  getAll(): T[] {
    return this.items;
  }

  filter(callback: (item: T) => boolean): T[] {
    return this.items.filter(callback);
  }
}

function renderProduct(p: Product): string {
  const createProductCard = (id: string, content: string, price: number): string => {
    return `
     <div class="item" id="${id}">
       ${content}
       <div class="price">$${price}</div>
     </div>
   `;
  };

  if (p.type === "book") {
    const content = `<strong>Book:</strong> ${p.title} by ${p.author}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "electronics") {
    const warranty = p.warranty ? ` - Warranty: ${p.warranty} year(s)` : "";
    const content = `<strong>Electronics:</strong> ${p.item} - ${p.model}${warranty}`;
    return createProductCard(p.id, content, p.price);
  }

  if (p.type === "clothing") {
    const size = p.size ? ` - Size ${p.size}` : "";
    const content = `<strong>Clothing:</strong> ${p.item} by ${p.brand}${size}`;
    return createProductCard(p.id, content, p.price);
  }

  const _never: never = p;
  throw new Error(`Unknown product type: ${JSON.stringify(p)}`);
}


const products = new Collection<Product>([
  { id: "c1", type: "clothing", item: "Jacket", brand: "Northloom", size: "M", price: 89.99 },
  { id: "e1", type: "electronics", item: "Tablet", model: "Pixelon Slate-A9", warranty: 2, price: 349.99 },
  { id: "b1", type: "book", title: "Dune", author: "Frank Herbert", price: 14.99 },

  { id: "e2", type: "electronics", item: "Smartphone", model: "NovaCore X1-Alpha", warranty: 2, price: 699.99 },
  { id: "c2", type: "clothing", item: "Hoodie", brand: "CozyForge", size: "S", price: 49.99 },
  { id: "b2", type: "book", title: "1984", author: "George Orwell", price: 9.99 },

  { id: "e3", type: "electronics", item: "Headphones", model: "EchoSphere Silent-7", price: 129.99 },
  { id: "c3", type: "clothing", item: "Jeans", brand: "BlueWeave", size: "L", price: 59.99 },
  { id: "e4", type: "electronics", item: "Laptop", model: "HexaBook Orion-15", warranty: 3, price: 1199.99 },

  { id: "b3", type: "book", title: "Brave New World", author: "Aldous Huxley", price: 11.99 },
  { id: "c4", type: "clothing", item: "T-Shirt", brand: "Fabricon", price: 19.99 },
  { id: "e5", type: "electronics", item: "Smartwatch", model: "Chronex Pulse-Q", warranty: 1, price: 199.99 },

  { id: "b4", type: "book", title: "The Lord of the Rings", author: "J.R.R. Tolkien", price: 24.99 },
]);

const output = document.getElementById("output");

// Track currently active filter category (or undefined for all)
let currentFilter: Product["type"] | undefined = undefined;

function showProducts(filter?: Product["type"]): void {
  if (!output) return;

  let itemsToShow: Product[];

  // Toggle logic
  if (currentFilter === filter || filter === undefined) {
    // Reset to all if same button clicked again or filter is undefined
    itemsToShow = products.getAll();
    currentFilter = undefined;
    setActiveButton(undefined);
  } else {
    itemsToShow = products.filter((p: Product): boolean => p.type === filter);
    currentFilter = filter;
    setActiveButton(filter);
  }

  output.innerHTML = itemsToShow.map(renderProduct).join("");
}

// Active Button Highlighting
function setActiveButton(filter: Product["type"] | undefined): void {
  const buttons: { [key: string]: HTMLElement | null } = {
    all: document.getElementById("all"),
    book: document.getElementById("books"),
    electronics: document.getElementById("electronics"),
    clothing: document.getElementById("clothing")
  };

  for (const key in buttons) {
    buttons[key]?.classList.remove("active");
  }

  if (filter === undefined) {
    buttons["all"]?.classList.add("active");
  } else {
    buttons[filter]?.classList.add("active");
  }
}

// Button Event Handlers
document.getElementById("all")?.addEventListener("click", (): void => showProducts());
document.getElementById("books")?.addEventListener("click", (): void => showProducts("book"));
document.getElementById("electronics")?.addEventListener("click", (): void => showProducts("electronics"));
document.getElementById("clothing")?.addEventListener("click", (): void => showProducts("clothing"));

// Initial Render
document.addEventListener("DOMContentLoaded", (): void => showProducts());
```
