---
id: 61329d52e5010e08d9b9d66b
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kipengele kingine muhimu cha `meta` kwa ajili ya Ufikikaji na SEO ni ufafanuzi wa `description`. Thamani ya sifa ya kipengele ya `content` hutumika na mashine za utafutaji kutoa maelezo ya ukurasa wako.

Ongeza kipengele cha `meta` chenye sifa ya kipengele ya `name` iliyowekwa kuwa `description`, na kipa sifa ya kipengele ya `content` yenye manufaa.

# --hints--

Unapaswa kuongeza kipengele kipya cha `meta` ndani ya `head`.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

Unapaswa kumpa `meta` sifa ya kipengele ya `name` yenye thamani ya `description`.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

Unapaswa kumpa `meta` sifa ya kipengele ya `content`.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

Thamani ya sifa ya kipengele ya `content` haipaswi kuzidi herufi 165. _Hii ni urefu wa maelezo wa juu kabisa wa Google._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
