---
id: 613e2546d0594208229ada50
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Ndani ya `header`, toa muktadha kuhusu ukurasa kwa kupanga ndani kipengele kimoja cha `img`, `h1`, na `nav`.

`img` inapaswa kuonyesha `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`, kuwa na `id` ya `logo`, na kuwa na maandishi ya `alt` ya `freeCodeCamp`.

`h1` inapaswa kuwa na maandishi ya `HTML/CSS Quiz`.

# --hints--

Unapaswa kuongeza kipengele cha `img` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > img'));
```

Unapaswa kuongeza kipengele cha `h1` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > h1'));
```

Unapaswa kuongeza kipengele cha `nav` kwenye kipengele cha `header`.

```js
assert.exists(document.querySelector('header > nav'));
```

Unapaswa kuweka vipengele vya `img`, `h1`, na `nav` kwa mpangilio huo.

```js
assert.exists(document.querySelector('img + h1 + nav'));
```

Unapaswa kumpa kipengele cha `img` sifa ya kipengele ya `src` yenye thamani ya `https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg`.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg');
```

Unapaswa kumpa kipengele cha `img` sifa ya kipengele ya `id` yenye thamani ya `logo`.

```js
assert.equal(document.querySelector('img')?.id, 'logo');
```

Unapaswa kumpa kipengele cha `img` sifa ya kipengele ya `alt` yenye thamani ya `freeCodeCamp`.

```js
assert.equal(document.querySelector('img')?.alt, 'freeCodeCamp');
```

Unapaswa kumpa kipengele cha `h1` maandishi ya `HTML/CSS Quiz`.

```js
assert.include(document.querySelector('h1')?.innerText?.toLowerCase(), 'html/css quiz');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
