---
id: 614ccc21ea91ef1736b9b578
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Karibu katika sehemu ya kwanza ya Jaribio fupi la Ufikikaji. Unapoanza kuwa msanidi programu mwenye uzoefu wa HTML na CSS, tumekuanzishia na muundo wa msingi.

Anza safari hii ya ufikikaji kwa kutoa sifa ya kipengele ya `lang` kwa kipengele chako cha `html`. Hii itasaidia kisomaji cha skrini kutambua lugha ya ukurasa.

# --hints--

Unapaswa kumpa kipengele cha `html` sifa ya kipengele ya `lang`. _Kidokezo: Unaweza kutumia thamani ya `en` kwa Kiingereza._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
