---
id: 67353ab907569125112caffa
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

이제 함수를 사용할 시간입니다. 네 개의 인수 `processIngredient`, `baseQuantity`, `baseUnit`, `newUnit`를 받는 새로운 `newServings` 함수를 선언하세요.

함수는 기본 수량을 새로운 인분 수에 맞게 조정하고, 그 조정된 수량을 기본 단위에서 새로운 단위로 변환한 뒤, 변환된 수량을 소수점 둘째 자리까지 고정된 문자열로 반환해야 합니다.

막히면 기록해둔 함수 호출 예제를 참고하세요. 두 예제의 값을 모두 변경해가며 실험할 수 있습니다.

# --hints--

`processIngredient` 함수를 선언해야 합니다.

```js
assert.isFunction(processIngredient);
```

`processIngredient` 함수는 `baseQuantity`을 첫 번째 매개변수로 가져야 합니다.

```js
assert.match(processIngredient?.toString(), /\(baseQuantity/);
```

`processIngredient` 함수는 `baseUnit`을 두 번째 매개변수로 가져야 합니다.

```js
assert.match(processIngredient?.toString(), /\(baseQuantity,\s*baseUnit/);
```

`processIngredient` 함수는 `newUnit`을 세 번째 매개변수로 가져야 합니다.

```js
assert.match(processIngredient?.toString(), /\(baseQuantity,\s*baseUnit,\s*newUnit/);
```

`processIngredient` 함수는 `newServings`을 네 번째 매개변수로 가져야 합니다.

```js
assert.match(processIngredient?.toString(), /\(baseQuantity,\s*baseUnit,\s*newUnit,\s*newServings\)/);
```

`processIngredient` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(processIngredient(10, "gram", "cup", 5));
```

`processIngredient` 함수는 올바른 값을 반환해야 합니다.

```js
assert.strictEqual(processIngredient(10, "gram", "cup", 5), "0.21");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Recipe Ingredient Converter</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main class="container">
      <h1>Recipe Ingredient Converter</h1>
      <form id="recipe-form">
        <div class="input-container">
          <label for="ingredient">Ingredient:</label>
          <input type="text" id="ingredient" required />
        </div>

        <div class="input-container">
          <label for="quantity">Quantity:</label>
          <input type="number" id="quantity" step="any" required />
        </div>

        <div class="input-container">
          <label for="unit">Unit:</label>
          <select id="unit">
            <option value="cup">Cup</option>
            <option value="gram">Gram</option>
            <option value="ounce">Ounce</option>
            <option value="teaspoon">Teaspoon</option>
          </select>
        </div>
        <div class="input-container">
          <label for="servings">Number of Servings:</label>
          <input type="number" id="servings" step="any" value="1" />
        </div>
        <button type="submit">Convert</button>
      </form>

      <section>
        <h2>Converted Ingredients</h2>
        <ul id="result-list"></ul>
      </section>
    </main>
    <script src="script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --dark-grey: #0a0a23;
  --white: #fff;
  --light-grey: #ddd;
  --golden-yellow: #fecc4c;
  --yellow: #ffcc4c;
  --gold: #feac32;
  --orange: #ffac33;
  --dark-orange: #f89808;
}

body {
  font-family: Arial, sans-serif;
  background-color: var(--dark-grey);
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
}

h1 {
  margin: 10px 0 15px;
}

.container {
  text-align: center;
  background: var(--white);
  padding: 20px;
  border-radius: 8px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
  width: 100%;
  max-width: 600px;
}

input,
select {
  padding: 10px;
  font-size: 16px;
  margin: 5px 0;
  border: 1px solid var(--light-grey);
  border-radius: 4px;
}

button {
  font-size: 1.2rem;
  cursor: pointer;
  width: 200px;
  margin: 20px 0 30px;
  color: var(--dark-grey);
  background-color: var(--gold);
  background-image: linear-gradient(var(--golden-yellow), var(--orange));
  border-color: var(--gold);
  border-width: 3px;
}

button:hover {
  background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

ul {
  list-style-type: none;
  padding: 0;
}

li {
  padding: 10px;
  border-bottom: 1px solid var(--light-grey);
  display: flex;
  justify-content: space-between;
}
```

```js
const conversionTable = {
  cup: { gram: 240, ounce: 8.0, teaspoon: 48 },
  gram: { cup: 1 / 240, ounce: 0.0353, teaspoon: 0.2 },
  ounce: { cup: 0.125, gram: 28.35, teaspoon: 6 },
  teaspoon: { cup: 1 / 48, gram: 5, ounce: 0.167 },
}

const convertQuantity = (fromUnit) => (toUnit) => (quantity) => {
  const conversionRate = conversionTable[fromUnit][toUnit];
  return quantity * conversionRate;
}

const gramsResult = convertQuantity("cup")("gram")(2);
console.log(gramsResult);

const adjustForServings = (baseQuantity) => (newServings) =>
  baseQuantity * newServings;

const servingsResult = adjustForServings(4)(6);
console.log(servingsResult);

--fcc-editable-region--

--fcc-editable-region--
```
