---
id: a3bfc1673c0526e06d3ac698
title: Construye una Calculadora de Suma de Números Primos
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

En este laboratorio, construirás una calculadora que toma un número y devuelve la suma de todos los números primos que son menores o iguales a ese número.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario:**

1. Debes tener una función `sumPrimes` que acepte un número como argumento.
1. La función `sumPrimes` debe devolver la suma de todos los números primos menores o iguales al número proporcionado.
1. Si el número de entrada es menor que `2`, la función debe devolver `0`.

# --hints--

Deberías tener una función `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` debería devolver `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` debería devolver `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` debería devolver `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` debería devolver `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` debería devolver `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
