---
id: 66d75dd0aa65a71600dc669b
title: بناء برنامج إدارة المخزون
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

في هذا المختبر، ستبني برنامجًا لإدارة المخزون يسمح لك بإضافة المنتجات وتحديثها والعثور عليها وإزالتها من المخزون. ستستخدم مصفوفة من الكائنات لتمثيل المخزون، حيث يحتوي كل كائن على مفتاحين هما name و quantity.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تعلن مصفوفة فارغة باسم `inventory` ستخزّن كائنات منتجات تحمل مفتاح `name` بقيمة سلسلة نصية بأحرف صغيرة، ومفتاح `quantity` بقيمة عدد صحيح.
1. يجب أن تنشئ دالة باسم `findProductIndex` تأخذ اسم المنتج كوسيطة لها وتُعيد فهرس كائن المنتج المقابل داخل مصفوفة `inventory`. وينبغي أن تستخدم الدالة دائمًا الصيغة ذات الأحرف الصغيرة لاسم المنتج لإجراء البحث. وإذا لم يُعثر على المنتج، فيجب أن تُعيد الدالة `-1`.
1. يجب أن تنشئ دالة باسم `addProduct` تأخذ كائن منتج كوسيطة لها.
1. إذا كان المنتج موجودًا بالفعل في المخزون، فيجب أن تحدّث الدالة `addProduct` قيمة `quantity` الخاصة به بإضافة الكمية المُمرَّرة إلى الدالة إلى `quantity` الحالية، وأن تسجّل في وحدة التحكم اسم المنتج متبوعًا بمسافة و`quantity updated`.
1. إذا لم يكن المنتج موجودًا في المخزون، فيجب أن تدفع الدالة `addProduct` المنتج إلى مصفوفة `inventory` وأن تسجّل اسم المنتج في وحدة التحكم، متبوعًا بمسافة و`added to inventory`.
1. يجب أن تنشئ دالة باسم `removeProduct` تأخذ اسم المنتج والكمية كوسيطتين لها.
1. يجب أن تطرح الدالة `removeProduct` الكمية المُمرَّرة من كائن المنتج المقابل داخل المخزون وأن تسجّل السلسلة `Remaining <product-name> pieces: <product-quantity>` في وحدة التحكم، حيث يُستبدل `<product-name>` باسم المنتج، ويُستبدل `<product-quantity>` بكمية المنتج.
1. إذا كانت الكمية بعد الطرح صفرًا، فيجب أن تزيل `removeProduct` كائن المنتج من المخزون. وإذا كانت الكمية في المخزون غير كافية لإجراء الطرح، فيجب أن تسجّل الدالة `removeProduct` السلسلة `Not enough <product-name> available, remaining pieces: <product-quantity>` في وحدة التحكم.
1. إذا لم يكن المنتج المراد إزالته موجودًا في المخزون، فيجب أن تسجّل الدالة `removeProduct` السلسلة `<product-name> not found` في وحدة التحكم.

**ملاحظة:** لتجنب التعارضات، احتفظ فقط بالتسجيلات المذكورة في قصص المستخدم عند تشغيل الاختبارات.

# --hints--

يجب إعلان مصفوفة فارغة باسم `inventory`.

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

يجب أن تحتوي على دالة باسم `findProductIndex`.

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` يجب أن تُرجع فهرس الكائن الذي يحتوي على `name` يساوي `"flour"` داخل مصفوفة `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` يجب أن تُرجع فهرس الكائن الذي يحتوي على `name` يساوي `"flour"` داخل مصفوفة `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

`findProductIndex("Flour")` يجب أن تُرجع `-1` عندما لا يُعثر على كائن يحتوي على `name` يساوي `"flour"` داخل مصفوفة `inventory`.

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

يجب أن تحتوي على دالة باسم `addProduct`.

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` يجب أن تضيف `5` إلى قيمة `quantity` للكائن الذي يحتوي على `name` يساوي `"flour"` في مصفوفة `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`addProduct({name: "FLOUR", quantity: 5})` يجب أن تسجل `flour quantity updated` عندما يُعثر على كائن يحتوي على `name` يساوي `"flour"` في مصفوفة `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

`addProduct({name: "FLOUR", quantity: 5})` يجب أن تضيف `{name: "flour", quantity: 5}` إلى مصفوفة `inventory` عندما لا يُعثر على كائن يحتوي على `name` يساوي `"flour"` في المخزون.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

`addProduct({name: "FLOUR", quantity: 5})` يجب أن تسجل `flour added to inventory` عندما لا يُعثر على كائن يحتوي على `name` يساوي `"flour"` في المخزون.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

يجب إنشاء دالة باسم `removeProduct`.

```js
assert.isFunction(removeProduct);
```

`removeProduct("FLOUR", 5)` يجب أن تسجل `flour not found` عندما لا يُعثر على كائن يحتوي على `name` يساوي `"flour"` في مصفوفة `inventory`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` يجب أن تخصم `5` من قيمة `quantity` للكائن الذي يحتوي على `name` يساوي `"flour"` داخل مصفوفة `inventory`.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` يجب أن تسجل `Remaining flour pieces: 15` في وحدة التحكم، عندما تكون قيمة `inventory` تساوي `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

إذا كانت قيمة `quantity` بعد الطرح صفرًا، يجب أن تزيل دالة `removeProduct` كائن المنتج من المخزون.

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` يجب أن تسجل `Not enough flour available, remaining pieces: 5` عندما تكون قيمة `inventory` تساوي `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]`.

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
