---
id: 61968f8877c6720d6d61aaf5
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Fügen Sie innerhalb des `body` ein `div` mit einem `class` von `ground` hinzu.

# --hints--

Sie sollten ein neues `div` hinzufügen.

```js
assert.exists(document.querySelector('div'));
```

Sie sollten dem `div` ein `class` von `ground` geben.

```js
assert.include(document.querySelector('div')?.className, 'ground');
```

Sie sollten das `div` innerhalb des `body` platzieren.

```js
assert.exists(document.querySelector('body > div.ground'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}
```
