---
id: 619665c9abd72906f3ad30f9
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Sie werden einen fröhlichen Flappy Penguin erstellen und dabei CSS-Transformationen und Animationen weiter erkunden.

Beginnen Sie mit Ihrem grundlegenden HTML-Boilerplate-Code. Fügen Sie die `DOCTYPE`-Deklaration, das `html`-Element mit auf Englisch eingestellter Sprache, die passenden `meta`-Tags, ein `head`, `body` und `title`-Element hinzu. Verlinken Sie außerdem Ihr Stylesheet mit der Seite.

# --hints--

Ihr Code sollte eine `<!DOCTYPE html>`-Deklaration enthalten.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Ihr Code sollte ein `html`-Element enthalten.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Ihr `html`-Element sollte ein öffnendes Tag mit einem `lang`-Attribut von `en` haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Ihr `html`-Element sollte ein schließendes Tag haben.

```js
assert(code.match(/<\/html\s*>/gi));
```

Ihr Code sollte ein `head`-Element innerhalb des `html`-Elements enthalten.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Ihr Code sollte ein `body`-Element innerhalb des `html`-Elements enthalten.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Ihr `head`-Element sollte vor Ihrem `body`-Element stehen.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Sie sollten zwei `meta`-Elemente haben.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Ein `meta`-Element sollte ein `name`-Attribut mit dem Wert `viewport` und `content` mit dem Wert `width=device-width, initial-scale=1.0` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

Das andere `meta`-Element sollte das `charset`-Attribut mit dem Wert `UTF-8` haben.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Ihr Code sollte ein `title`-Element enthalten.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Ihr `title` sollte etwas Text enthalten.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Ihr Code sollte ein `link`-Element enthalten.

```js
assert.exists(document.querySelector('link'));
```

Ihr `link`-Element sollte innerhalb Ihres `head`-Elements sein.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

Ihr `link`-Element sollte ein `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Ihr `link`-Element sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
