---
id: 6842a6cd9836f0114a5b7a8a
title: بناء مولد نمط أرقام
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

في هذا المختبر ستتدرب على أساسيات Python ببناء تطبيق صغير يُنشئ نمط أرقام.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `number_pattern` تأخذ معلمة واحدة `n` (تمثل عددًا صحيحًا موجبًا).
1. يجب أن تستخدم `number_pattern` حلقة `for`.
1. يجب أن تُرجع `number_pattern(n)` سلسلة نصية تحتوي كل الأعداد الصحيحة بدءًا من 1 حتى `n` (شاملة) مفصولة بمسافة. على سبيل المثال، يجب أن تُرجع `number_pattern(4)` السلسلة النصية `1 2 3 4`.
1. إذا كانت المعلمة الممررة إلى الدالة ليست قيمة عدد صحيح، يجب أن تُرجع الدالة `Argument must be an integer value.`
1. إذا كانت المعلمة الممررة إلى الدالة أقل من 1، يجب أن تُرجع الدالة `Argument must be an integer greater than 0.`

# --hints--

يجب أن يكون لديك دالة `number_pattern`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

يجب أن تحتوي دالة `number_pattern` على معلمة واحدة باسم `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

يجب أن تُرجع `number_pattern(4)` القيمة `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

يجب أن تُرجع `number_pattern(12)` القيمة `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

يجب أن تُرجع `number_pattern` قائمة أرقام مفصولة بمسافة لأي عدد صحيح موجب.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2
    
    `)
})
```

يجب أن تُرجع `number_pattern` القيمة `Argument must be an integer value.` عند تمرير قيمة ليست عددًا صحيحًا.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

يجب أن تُرجع `number_pattern` القيمة `Argument must be an integer greater than 0.` عند تمرير عدد صحيح غير موجب.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
