---
id: 66c06d618d075c7f7f1b890a
title: 운세 점쟁이 만들기
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

이 실습에서는 다섯 가지 운세를 정의하고 그중 하나를 무작위로 선택해 사용자에게 보여줍니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. `fortune1`, `fortune2`, `fortune3`, `fortune4`, `fortune5` 다섯 변수를 원하는 문자열 값으로 초기화하세요. 아래 옵션을 사용해도 됩니다:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. 1부터 5까지(포함) 무작위 숫자를 선택해 `randomNumber` 변수에 할당하세요.

3. `selectedFortune` 변수를 만들고 다음 규칙에 따라 적절한 운세를 할당하세요:

   - `randomNumber`가 1이면 `fortune1` 값을 `selectedFortune`에 할당하세요.
   - `randomNumber`가 2이면 `fortune2` 값을 `selectedFortune`에 할당하세요.
   - `randomNumber`가 3이면 `fortune3` 값을 `selectedFortune`에 할당하세요.
   - `randomNumber`가 4이면 `fortune4` 값을 `selectedFortune`에 할당하세요.
   - `randomNumber`가 5이면 `fortune5` 값을 `selectedFortune`에 할당하세요.

4. `selectedFortune`을 콘솔에 출력하세요.

# --hints--

`fortune1`을 문자열 값으로 초기화하세요.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

`fortune2`을 문자열 값으로 초기화하세요.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

`fortune3`을 문자열 값으로 초기화하세요.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

`fortune4`을 문자열 값으로 초기화하세요.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

`fortune5`을 문자열 값으로 초기화하세요.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

무작위 숫자를 생성하려면 `Math.random()` 메서드를 사용하세요.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

1부터 5까지(포함) 무작위 숫자를 생성해 `randomNumber` 변수에 할당하세요.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

`selectedFortune` 값에 따라 값을 할당받는 `randomNumber` 변수를 만드세요.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber`은 해당 운세와 일치해야 합니다. 예를 들어 `randomNumber`이 1이면 `selectedFortune`는 `fortune1`과 같아야 합니다.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

`selectedFortune`을 콘솔에 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
