---
id: 66cc095dfe1682753d2ab030
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Uno dei concetti che hai imparato nelle lezioni precedenti è stato come restituire valori da una funzione.

Ecco un promemoria su come restituire un valore da una funzione:

```js
function myFunction() {
  return "Hello World";
}
```

All'interno della tua funzione `addTwoAndSeven`, restituisci la somma di `2` e `7`.

# --hints--

La tua funzione `addTwoAndSeven` dovrebbe restituire la somma di `2` e `7`.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

La tua funzione `addTwoAndSeven` dovrebbe restituire la somma usando l'operatore `+`.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
