---
id: 66cc1a3a39aef47d6473cb2f
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Il passaggio successivo è aggiungere la funzionalità di divisione alla tua calcolatrice.

Dichiara una funzione chiamata `calculateQuotient` che prende due parametri, `num1` e `num2`.

All'interno della funzione, restituisci il quoziente dei due parametri.

Poi, aggiungi una `console.log` che chiama la funzione `calculateQuotient` con gli argomenti `7` e `11`.

# --hints--

Dovresti avere una funzione chiamata `calculateQuotient`.

```js
assert.isFunction(calculateQuotient);
```

La tua funzione `calculateQuotient` dovrebbe avere due parametri chiamati `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateQuotient', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

La tua funzione `calculateQuotient` dovrebbe restituire il quoziente dei due parametri.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Dovresti avere una `console.log` che chiama la funzione `calculateQuotient` con gli argomenti `7` e `11`.

```js
assert.match(code, /console\.log\s*\(\s*calculateQuotient\s*\(\s*7\s*,\s*11\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
