---
id: 66cc0f1ae40802781b2ea972
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Nelle lezioni precedenti, hai imparato come lavorare con i parametri delle funzioni.

Un parametro di una funzione è una variabile definita nella dichiarazione della funzione e funge da segnaposto.

Ecco un esempio di una funzione che ha un parametro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Dichiara una funzione chiamata `calculateSum` che prende due parametri, `num1` e `num2`.

# --hints--

Dovresti dichiarare una funzione chiamata `calculateSum`.

```js
assert.isFunction(calculateSum);
```

La tua funzione `calculateSum` dovrebbe prendere due parametri, `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateSum', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
