---
id: 66cc0de2920bec775f610424
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Anche se le funzioni funzionano come previsto, nel tuo codice cՏ molta ripetizione.

Dato che stai costruendo una calcolatrice, non vuoi dover creare una funzione per ogni possibile combinazione di numeri da sommare.

Sarebbe meglio creare una singola funzione riutilizzabile che possa sommare insieme due numeri qualsiasi.

Rimuovi tutto il codice che hai scritto finora.

Nel passaggio successivo, vedrai come lavorare con parametri e argomenti nelle funzioni.

# --hints--

Non dovresti avere una funzione chiamata `addTwoAndSeven`.

```js
const explorer = await __helpers.Explorer(code);
const { addTwoAndSeven } = explorer.allFunctions;
assert.notExists(addTwoAndSeven);
```

Non dovresti avere una chiamata a funzione a `addTwoAndSeven`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addTwoAndSeven())";
assert.notInclude(strippedCode, unexpectedStatement);
```

Non dovresti avere una funzione chiamata `addThreeAndFour`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
assert.notExists(addThreeAndFour);
```

Non dovresti avere una chiamata a funzione a `addThreeAndFour`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const unexpectedStatement = "console.log(addThreeAndFour())";
assert.notInclude(strippedCode, unexpectedStatement);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
