---
id: 66cc0ba5881acb7692cfc4de
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora ├© il momento di aggiungere un'altra funzione.

Dichiara una funzione chiamata `addThreeAndFour` che restituisce la somma di `3` e `4`.

Poi chiama la funzione `addThreeAndFour` all'interno di un `console.log` per vedere il risultato.

# --hints--

Dovresti dichiarare una funzione chiamata `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

La tua funzione `addThreeAndFour` dovrebbe restituire la somma di `3` e `4`.

```js
assert.equal(addThreeAndFour(), 7);
```

La tua funzione `addThreeAndFour` dovrebbe restituire la somma usando l'operatore `+`.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Dovresti chiamare la funzione `addThreeAndFour` all'interno di un `console.log`.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
