---
id: 66cc14ca14e65e7a9ebecb80
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Ora che la tua calcolatrice pu├▓ sommare due numeri, ├© il momento di creare una funzione che sottragga due numeri.

Dichiara una funzione chiamata `calculateDifference` che prende due parametri, `num1` e `num2`.

All'interno della funzione, restituisci la differenza tra i due parametri.

# --hints--

Dovresti avere una funzione chiamata `calculateDifference`.

```js
assert.isFunction(calculateDifference);
```

La tua funzione `calculateDifference` dovrebbe avere due parametri chiamati `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

La tua funzione `calculateDifference` dovrebbe restituire la differenza tra i due parametri.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
