---
id: 66cc2739f687e484d50bb6f1
title: Passaggio 18
challengeType: 1
dashedName: step-18
---

# --description--

L'ultima funzionalità per la tua calcolatrice sarà la radice quadrata.

Un quadrato è un numero moltiplicato per se stesso. Quindi, la radice quadrata è il numero che, moltiplicato per se stesso, dà il numero originale.

Per esempio, la radice quadrata di `9` è `3` perché `3 * 3 = 9`.

Per ottenere la radice quadrata di un numero in JavaScript, puoi usare il metodo <dfn>Math.sqrt()</dfn>.

Ecco un esempio:

```js
// result: 3
Math.sqrt(9);
```

Dichiara una funzione chiamata `calculateSquareRoot` che prende un parametro `num` e restituisce la radice quadrata di `num`.

# --hints--

Dovresti avere una funzione chiamata `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

La tua funzione `calculateSquareRoot` dovrebbe prendere un parametro `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

La tua funzione `calculateSquareRoot` dovrebbe restituire la radice quadrata di `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
