---
id: 66cc0f1ae40802781b2ea972
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Nelle lezioni precedenti, hai imparato come lavorare con i parametri delle funzioni.

Un parametro di una funzione è una variabile definita nella dichiarazione della funzione e funge da segnaposto.

Ecco un esempio di una funzione che ha un parametro:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Dichiara una funzione chiamata `calculateSum` che prende due parametri, `num1` e `num2`.

# --hints--

Dovresti dichiarare una funzione chiamata `calculateSum`.

```js
assert.isFunction(calculateSum);
```

La tua funzione `calculateSum` dovrebbe prendere due parametri, `num1` e `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
