---
id: 66cc0ba5881acb7692cfc4de
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora è il momento di aggiungere un'altra funzione.

Dichiara una funzione chiamata `addThreeAndFour` che restituisce la somma di `3` e `4`.

Poi chiama la funzione `addThreeAndFour` all'interno di un `console.log` per vedere il risultato.

# --hints--

Dovresti dichiarare una funzione chiamata `addThreeAndFour`.

```js
assert.isFunction(addThreeAndFour);
```

La tua funzione `addThreeAndFour` dovrebbe restituire la somma di `3` e `4`.

```js
assert.strictEqual(addThreeAndFour(), 7);
```

La tua funzione `addThreeAndFour` dovrebbe restituire la somma usando l'operatore `+`.

```js
const explorer = await __helpers.Explorer(code);
const { addThreeAndFour } = explorer.allFunctions;
const validExpressions = ["3 + 4", "4 + 3", "(3 + 4)", "(4 + 3)"];
assert.isTrue(validExpressions.some(expr => addThreeAndFour?.hasReturn(expr)));
```

Dovresti chiamare la funzione `addThreeAndFour` all'interno di un `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addThreeAndFour())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
