---
id: 66cc0a9e06e00b75d6782be9
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Nelle lezioni precedenti, hai imparato come chiamare (invocare) una funzione. Chiamare una funzione significa eseguire il codice all'interno della funzione.

Ecco un promemoria su come chiamare una funzione:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Aggiungi un'istruzione `console.log` e al suo interno chiama la funzione `addTwoAndSeven`.

Ora dovresti vedere la somma di `2` e `7` stampata nella console.

# --hints--

Dovresti chiamare la funzione `addTwoAndSeven` all'interno di una `console.log`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
