---
id: 66cc12fa504b0479dac479a0
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Il vantaggio delle funzioni è che possono essere chiamate con argomenti diversi, permettendoti di riutilizzare lo stesso codice con vari valori.

Ecco alcuni esempi di chiamate a funzioni con argomenti stringa diversi:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Aggiungi un `console.log` che chiami la funzione `calculateSum` con gli argomenti `2` e `5`.

# --hints--

Dovresti avere un `console.log` che chiami la funzione `calculateSum` con gli argomenti `2` e `5`.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
