---
id: 66cc06012e46aa6bc9b8c001
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, rivedrai come lavorare con le funzioni costruendo un'app calcolatrice.

Nelle lezioni precedenti, hai imparato a dichiarare funzioni in questo modo:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Inizia creando una funzione chiamata `addTwoAndSeven`.

Puoi scegliere di usare la sintassi della funzione tradizionale o quella della funzione freccia.

# --hints--

Dovresti dichiarare una funzione chiamata `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
