---
id: 66cc10958d72e57907f898b2
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

All'interno della tua funzione `calculateSum`, dovrai restituire la somma dei due parametri.

# --hints--

La tua funzione `calculateSum` dovrebbe restituire la somma dei due parametri.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
