---
id: 66cc16abc30b367ba1d1f9ef
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Ora è il momento di testare la tua funzione `calculateDifference`.

Inizia aggiungendo una `console.log` che chiama la funzione `calculateDifference` con gli argomenti `22` e `5`.

Poi, aggiungi una seconda `console.log` che chiama la funzione `calculateDifference` con gli argomenti `12` e `1`.

Infine, aggiungi una terza `console.log` che chiama la funzione `calculateDifference` con gli argomenti `17` e `9`.

# --hints--

Dovresti avere una `console.log` che chiama la funzione `calculateDifference` con gli argomenti `22` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(22,5))";
assert.include(strippedCode, expectedStatement);
```

Dovresti avere una `console.log` che chiama la funzione `calculateDifference` con gli argomenti `12` e `1`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(12,1))";
assert.include(strippedCode, expectedStatement);
```

Dovresti avere una `console.log` che chiama la funzione `calculateDifference` con gli argomenti `17` e `9`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateDifference(17,9))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
