---
id: 66cc0a9e06e00b75d6782be9
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Nelle lezioni precedenti, hai imparato come chiamare (invocare) una funzione. Chiamare una funzione significa eseguire il codice all'interno della funzione.

Ecco un promemoria su come chiamare una funzione:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Aggiungi un'istruzione `console.log` e al suo interno chiama la funzione `addTwoAndSeven`.

Ora dovresti vedere la somma di `2` e `7` stampata nella console.

# --hints--

Dovresti avere un `console.log` nel tuo codice.

```js
assert.match(code, /console\.log(.*)/);
```

Dovresti chiamare la funzione `addTwoAndSeven` all'interno del `console.log`.

```js
assert.match(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}
--fcc-editable-region--

--fcc-editable-region--
```
