---
id: 66cc172af7159a7c67804544
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Ora è il momento di aggiungere la funzionalità di moltiplicazione alla tua calcolatrice.

Dichiara una funzione chiamata `calculateProduct` che prende due parametri, `num1` e `num2`.

All'interno della funzione, restituisci il prodotto dei due parametri.

Poi, aggiungi una `console.log` che chiama la funzione `calculateProduct` con gli argomenti `13` e `5`.

# --hints--

Dovresti avere una funzione chiamata `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

La tua funzione `calculateProduct` dovrebbe avere due parametri chiamati `num1` e `num2`.

```js
const regex = __helpers.functionRegex('calculateProduct', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

La tua funzione `calculateProduct` dovrebbe restituire il prodotto dei due parametri.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Dovresti avere una `console.log` che chiama la funzione `calculateProduct` con gli argomenti `13` e `5`.

```js
assert.match(code, /console\.log\s*\(\s*calculateProduct\s*\(\s*13\s*,\s*5\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
