---
id: 66cc172af7159a7c67804544
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

Ora è il momento di aggiungere la funzionalità di moltiplicazione alla tua calcolatrice.

Dichiara una funzione chiamata `calculateProduct` che prende due parametri, `num1` e `num2`.

All'interno della funzione, restituisci il prodotto dei due parametri.

Poi, aggiungi una `console.log` che chiama la funzione `calculateProduct` con gli argomenti `13` e `5`.

# --hints--

Dovresti dichiarare una funzione chiamata `calculateProduct`.

```js
assert.isFunction(calculateProduct);
```

La tua funzione `calculateProduct` dovrebbe avere due parametri chiamati `num1` e `num2`.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

La tua funzione `calculateProduct` dovrebbe restituire il prodotto dei due parametri.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Dovresti avere una `console.log` che chiama la funzione `calculateProduct` con gli argomenti `13` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
