---
id: 66cc0de2920bec775f610424
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Anche se le funzioni funzionano come previsto, nel tuo codice cՏ molta ripetizione.

Dato che stai costruendo una calcolatrice, non vuoi dover creare una funzione per ogni possibile combinazione di numeri da sommare.

Sarebbe meglio creare una singola funzione riutilizzabile che possa sommare insieme due numeri qualsiasi.

Rimuovi tutto il codice che hai scritto finora.

Nel passaggio successivo, vedrai come lavorare con parametri e argomenti nelle funzioni.

# --hints--

Non dovresti avere una funzione chiamata `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Non dovresti avere una chiamata a funzione a `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Non dovresti avere una funzione chiamata `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Non dovresti avere una chiamata a funzione a `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
