---
id: 66cc26907ff6908402af0149
title: Passaggio 17
challengeType: 1
dashedName: step-17
---

# --description--

Per testare la tua funzione `calculateSquare`, dovrai chiamare la funzione un paio di volte.

Aggiungi una `console.log` che chiami la funzione `calculateSquare` con l'argomento `2`.

Poi, aggiungi un'altra `console.log` che chiami la funzione `calculateSquare` con l'argomento `9`.

# --hints--

Dovresti avere una `console.log` che chiami la funzione `calculateSquare` con l'argomento `2`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

Dovresti avere una `console.log` che chiami la funzione `calculateSquare` con l'argomento `9`.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
