---
id: 66cc21d23238dc8240a8a182
title: Passaggio 16
challengeType: 1
dashedName: step-16
---

# --description--

Sarebbe utile che la tua calcolatrice calcolasse il quadrato di un numero. Il quadrato di un numero è il numero moltiplicato per se stesso.

Per calcolare il quadrato di un numero in JavaScript, puoi usare il metodo <dfn>Math.pow()</dfn>. Il metodo `Math.pow()` prende due argomenti: il numero base e l'esponente.

Puoi anche usare l'operatore di esponenziazione (`**`) per calcolare il quadrato di un numero.

Ecco un esempio:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Dichiara una funzione chiamata `calculateSquare` che prende un parametro `num` e restituisce il quadrato di `num`.

# --hints--

Dovresti avere una funzione chiamata `calculateSquare`.

```js
assert.isFunction(calculateSquare);
```

La tua funzione `calculateSquare` dovrebbe prendere un parametro `num`.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

La tua funzione `calculateSquare` dovrebbe restituire il quadrato di `num`.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
