---
id: 66cdf76685e4cb5a8726e27b
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Per testare ulteriormente la tua funzione `calculateSum`, dovrai chiamarla con argomenti diversi.

Aggiungi un altro `console.log` che chiama la funzione `calculateSum` con gli argomenti `10` e `10`.

Sotto quel `console.log`, aggiungi un altro `console.log` che chiama la funzione `calculateSum` con gli argomenti `5` e `5`.

# --hints--

Dovresti avere un `console.log` che chiama la funzione `calculateSum` con gli argomenti `10` e `10`.

```js
assert.match(code, /console\.log\(calculateSum\(10\s*,\s*10\)\)/);
```

Dovresti avere un `console.log` che chiama la funzione `calculateSum` con gli argomenti `5` e `5`.

```js
assert.match(code, /console\.log\(calculateSum\(5\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

console.log(calculateSum(2, 5));

--fcc-editable-region--

--fcc-editable-region--
```
