---
id: 66cdf76685e4cb5a8726e27b
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Per testare ulteriormente la tua funzione `calculateSum`, dovrai chiamarla con argomenti diversi.

Aggiungi un altro `console.log` che chiama la funzione `calculateSum` con gli argomenti `10` e `10`.

Sotto quel `console.log`, aggiungi un altro `console.log` che chiama la funzione `calculateSum` con gli argomenti `5` e `5`.

# --hints--

Dovresti avere un `console.log` che chiama la funzione `calculateSum` con gli argomenti `10` e `10`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(10,10))";
assert.include(strippedCode, expectedStatement);
```

Dovresti avere un `console.log` che chiama la funzione `calculateSum` con gli argomenti `5` e `5`.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(5,5))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
