---
id: a97fd23d9b809dac9921074f
title: Побудувати функцію суми з необов’язковими аргументами
challengeType: 26
dashedName: build-an-optional-arguments-sum-function
---

# --description--

У цій лабораторній роботі ви створите функцію, яка приймає до двох аргументів і сумує їх, але якщо передано лише один аргумент, повертає функцію, яка чекає на друге число для суми.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Створіть функцію `addTogether`.
2. Якщо функція отримує два аргументи, вона має повернути суму цих двох аргументів.
3. Якщо передано лише один аргумент, `addTogether` має повернути функцію.
   - Коли викликають повернену функцію з одним аргументом, вона має повернути суму.

   ```js
   const sumTwoAnd = addTogether(2);
   sumTwoAnd(3); // 5
   ```

4. Якщо будь-який аргумент не є дійсним числом, поверніть `undefined`.

# --hints--

У вас має бути функція `addTogether`.

```js
assert.isFunction(addTogether);
```

`addTogether(2, 3)` має повертати 5.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)` має повертати 53.4.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)` має повертати `undefined`.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)` має повертати `undefined`.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")` має повертати `undefined`.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)` має повертати функцію.

```js
assert.isFunction(addTogether(5));
```

`addTogether(5)(7)` має повертати 12.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])` має повертати `undefined`.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")` має повертати `undefined`.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
