---
id: 61fd67a656743144844941cb
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

화면 읽기 프로그램은 문서 흐름에 따라 HTML 요소를 안내합니다. 결국 대차대조표에는 "Balance Sheet"라는 제목과 "AcmeWidgetCorp"라는 부제목이 있어야 합니다. 하지만 화면 읽기 프로그램이 안내할 때 이 순서는 적절하지 않습니다.

기존 `span`에 `class` 속성을 `flex`로 설정하고, 그 안에 두 개의 `span` 요소를 추가하세요. 첫 번째에는 텍스트 `AcmeWidgetCorp`를 넣으세요. 두 번째에는 텍스트 `Balance Sheet`를 넣으세요. CSS를 사용해 페이지에서 텍스트 순서를 뒤집겠지만, HTML 순서는 화면 읽기 프로그램에 더 적합할 것입니다.

# --hints--

기존 `span` 요소에 `class` 속성을 `flex`로 설정하세요.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

기존 `span` 요소 안에 두 개의 새로운 `span` 요소를 추가하세요.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

새로운 `span` 요소에는 `class` 속성이 없어야 합니다.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

첫 번째 새로운 `span` 요소에는 텍스트 `AcmeWidgetCorp`를 넣으세요.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent.trim(),
  'AcmeWidgetCorp'
);
```

두 번째 새로운 `span` 요소에는 텍스트 `Balance Sheet`를 넣으세요.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent.trim(),
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
