---
id: 61fd71d596e8f253b9408b39
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`tr` 요소는 표 행을 나타내는 데 사용됩니다. `tr` 요소 안에 `thead` 요소를 추가하세요. 새로 만든 `tr` 요소 안에 `td` 요소를 추가한 다음, `th` 요소 세 개를 추가하세요.

`td` 요소는 데이터 셀을 나타내고, `th` 요소는 헤더 셀을 나타냅니다.

# --hints--

`thead` 요소에는 `tr` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('thead')?.children?.[0]?.localName,
  'tr'
);
```

`tr` 요소에는 첫 번째 자식으로 `td` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[0]?.localName,
  'td'
);
```

`tr` 요소에는 `th` 요소 뒤에 세 개의 `td` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('tr')?.children?.[1]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[2]?.localName,
  'th'
);
assert.strictEqual(
  document.querySelector('tr')?.children?.[3]?.localName,
  'th'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>Assets</caption>
--fcc-editable-region--
            <thead>
            </thead>
--fcc-editable-region--
            <tbody></tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
