---
id: 61fd719788899952e67692b9
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

`thead`와 `tbody` 요소는 표에서 어느 부분이 헤더이고, 어느 부분이 주요 데이터나 내용인지 표시하는 데 사용됩니다.

첫 번째 `thead`에 `tbody`와 `table`를 `caption` 요소 아래에 추가하세요.

# --hints--

첫 번째 `table` 요소에는 `thead` 요소가 있어야 합니다.

```js
const table = document.querySelectorAll('table')?.[0];
assert.isNotNull(table?.querySelector('thead'));
```

첫 번째 `table` 요소에는 `tbody` 요소가 있어야 합니다.

```js
const table = document.querySelectorAll('table')?.[0];
const tbody = table?.querySelector('tbody');
assert.isNotNull(tbody);
```

`thead` 요소는 `caption` 요소 바로 아래에 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('caption')?.nextElementSibling?.localName,
  'thead'
);
```

`tbody` 요소는 `thead` 요소 바로 아래에 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('thead')?.nextElementSibling?.localName,
  'tbody'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
            <caption>Assets</caption>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
