---
id: 61fd77f7ad2aeb5ae34d07d6
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

두 번째 `tr` 요소에 텍스트 `th`를 가진 `Checking Our primary transactional account.` 요소를 추가하세요. `Checking `를 제외한 그 텍스트를 `span` 속성이 `class`로 설정된 `description` 요소로 감싸세요.

그 다음에, 다음 텍스트를 가진 세 개의 `td` 요소를 순서대로 추가하세요: `$54`, `$56`, `$53`. 세 번째 `td` 요소에 `class` 속성을 `current`로 설정하세요.

# --hints--

두 번째 `tr`에는 `th` 요소가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isNotNull(tableRow?.querySelector('th'));
```

`th` 요소에는 텍스트 `Checking Our primary transactional account.`가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th')?.innerText.trim(),
  'Checking Our primary transactional account.'
);
```

텍스트 `Our primary transactional account.`를 `span` 요소로 감싸야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(
  tableRow?.querySelector('th > span')?.textContent.trim(),
  'Our primary transactional account.'
);
```

`span` 요소에는 `class` 속성이 `description`로 설정되어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelector('th > span')?.classList?.contains('description')
);
```

세 개의 `td` 요소가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.lengthOf(tableRow?.querySelectorAll('td'), 3);
```

첫 번째 `td` 요소에는 텍스트 `$54`가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[0]?.textContent.trim(), '$54');
```

두 번째 `td` 요소에는 텍스트 `$56`가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[1]?.textContent.trim(), '$56');
```

세 번째 `td` 요소에는 텍스트 `$53`가 있어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.strictEqual(tableRow?.querySelectorAll('td')?.[2]?.textContent.trim(), '$53');
```

세 번째 `td` 요소에는 `class`가 `current`로 설정되어야 합니다.

```js
const tableBody = document.querySelector('tbody');
const tableRow = tableBody?.querySelectorAll('tr')?.[1];
assert.isTrue(
  tableRow?.querySelectorAll('td')?.[2]?.classList?.contains('current')
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
            <tbody>
              <tr class="data">
                <th>Cash <span class="description">This is the cash we currently have on hand.</span></th>
                <td>$25</td>
                <td>$30</td>
                <td class="current">$28</td>
              </tr>
--fcc-editable-region--
              <tr class="data">
              </tr>
--fcc-editable-region--
              <tr class="data"></tr>
              <tr class="total"></tr>
            </tbody>
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
