---
id: 61fd5a93fd62bb35968adeab
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

프로젝트를 시작할 때 `head` 요소에 스타일시트를 위한 `link` 요소를 추가하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.match(code, /<link/);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

`link` 요소는 값이 `rel`인 `stylesheet` 속성을 가져야 합니다.

```js
assert.match(code, /<link[\s\S]*?rel=('|"|`)stylesheet\1/);
```

`link` 요소는 값이 `href`인 `styles.css` 속성을 가져야 합니다.

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
