---
id: 61fd66c687e610436494c6f1
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

`section` 요소 안에 `h1` 요소를 추가하고, 그 안에 중첩된 `span` 요소를 넣으세요.

# --hints--

`section` 요소는 `h1` 요소를 포함해야 합니다.

```js
assert.strictEqual(
  document.querySelector('section')?.children?.[0]?.localName,
  'h1'
);
```

`h1` 요소는 `span` 요소를 포함해야 합니다.

```js
assert.strictEqual(
  document.querySelector('h1')?.children?.[0]?.localName,
  'span'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```

```css

```
