---
id: 61fd70336ebb3e4f62ee81ba
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML 표는 `caption` 요소를 사용해 표가 무엇에 관한 것인지 설명합니다. `caption` 요소는 항상 `table`의 첫 번째 자식이어야 하지만, `caption-side` CSS 속성으로 위치를 조정할 수 있습니다.

첫 번째 `caption`에 `table` 요소를 추가하고, 텍스트로 `Assets`을 넣으세요.

# --hints--

첫 번째 `table` 요소에는 `caption` 요소가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

`caption` 요소에는 텍스트 `Assets`이 있어야 합니다.

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
