---
id: 61fd6b7c83dbf54a08cf0498
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`div` 요소 안에 `span` 요소 세 개를 추가하세요. 각각에 `class` 속성을 `year`로 설정하고, 다음 텍스트를 순서대로 추가하세요: `2019`, `2020`, `2021`.

# --hints--

`div` 요소에는 `span` 요소 세 개가 있어야 합니다.

```js
assert.lengthOf(document.querySelector('div')?.children, 3);
```

각 `span` 요소에는 `class` 속성이 `year`로 설정되어야 합니다.

```js
const spans = [...document.querySelector('div')?.children];
spans.forEach(span => assert.isTrue(span?.classList?.contains('year')));
```

첫 번째 `span`에는 텍스트 `2019`가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[0]?.textContent.trim(),
  '2019'
);
```

두 번째 `span`에는 텍스트 `2020`가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[1]?.textContent.trim(),
  '2020'
);
```

세 번째 `span`에는 텍스트 `2021`가 있어야 합니다.

```js
assert.strictEqual(
  document.querySelector('div')?.children?.[2]?.textContent.trim(),
  '2021'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--
        <div id="years" aria-hidden="true">
        </div>
--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
