---
id: 699d20623e57e8ffee885454
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

目録カードを解析する最初のステップは、4つのフィールドを分けることです。`.split()` メソッドは、指定した区切り文字が現れるたびに文字列を配列の部分文字列に分割します。

例えば：

```js
"The Shining | King, Stephen | 1977 | Shelf K1".split("|")
// ["The Shining ", " King, Stephen ", " 1977 ", " Shelf K1"]
```

`parseCard` の中で、`.split("|")` に対して `rawString` を呼び出し、その結果を返してください。このステップの後、コンソールには4つの文字列からなる配列が表示されるはずです。

**注意：** 値の周りに余分な空白があるかもしれませんが、それは次のステップで整理します。

# --hints--

`.split("|")` 関数の中で `rawString` を `parseCard` に対して呼び出す必要があります。

```js
assert.match(parseCard.toString(), /rawString\s*\.\s*split\s*\(\s*('|"|`)\|\1\s*\)\s*;?/);
```

`parseCard` は配列を返すべきです。

```js
assert.isArray(parseCard("A | B | C | D"));
```

`"A | B | C | D"` は `parseCard` を引数に呼び出すと、4つの要素を持つ配列を返すべきです。

```js
assert.strictEqual(parseCard("A | B | C | D").length, 4);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
--fcc-editable-region--
  
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
