---
id: 699d20623e57e8ffee88545d
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

`"Unknown"` のチェックの後に、数値の年を持つ本の年代別グループ化のロジックを追加してください。`const decade = Math.floor(book.year / 10) * 10` を宣言します。例えば、`Math.floor(1977 / 10)` は `197` を返し、`10` を掛けると `1970` になります。

次に、`const decadeKey =` を `` `${decade}s` `` to produce a string like `"1970s"` と宣言してください。

ブラケット記法を使って本をグループ化します：`grouped[decadeKey]` がまだ存在しないかどうかを確認し、存在しなければ空の配列として初期化し、そこに `book` を追加してください。キーが変数に格納されているため、ここではブラケット記法が必要です。

# --hints--

`decade` という名前の変数を `Math.floor(book.year / 10) * 10` に設定して宣言してください。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decade\s*=\s*Math\.floor\s*\(\s*book\.year\s*\/\s*10\s*\)\s*\*\s*10/);
```

`decadeKey` という名前の変数を宣言し、`decade` と文字 `"s"` を含めてください。

```js
assert.match(groupByDecade.toString(), /(const|let|var)\s+decadeKey\s*=.*decade/s);
assert.match(groupByDecade.toString(), /decadeKey.*s/s);
```

`grouped[decadeKey]` がまだ存在しないかどうかを確認し、存在しなければ空の配列として初期化してください。

```js
assert.match(groupByDecade.toString(), /!\s*grouped\s*\[\s*decadeKey\s*\]/);
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*=\s*\[\s*\]/);
```

`book` を `grouped[decadeKey]` に追加してください。

```js
assert.match(groupByDecade.toString(), /grouped\s*\[\s*decadeKey\s*\]\s*\.push\s*\(\s*book\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
  for (let i = 0; i < rawCards.length; i++) {
    catalog.push(parseCard(rawCards[i]));
  }
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);

function findByAuthor(catalog, author) {
  const searchTerm = author.toLowerCase();
  const results = [];
  for (let i = 0; i < catalog.length; i++) {
    if (catalog[i].author.toLowerCase().includes(searchTerm)) {
      results.push(catalog[i]);
    }
  }
  return results;
}

function groupByDecade(catalog) {
  const grouped = {};
  for (let i = 0; i < catalog.length; i++) {
    const book = catalog[i];
    if (book.year === "Unknown") {
      if (!grouped["Unknown"]) {
        grouped["Unknown"] = [];
      }
      grouped["Unknown"].push(book);
      continue;
    }
--fcc-editable-region--

--fcc-editable-region--
  }
}
```
