---
id: 699d20623e57e8ffee885456
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

パーツが整理できたので、それらを構造化されたオブジェクトに変換できます。JavaScriptでは、オブジェクトは波括弧で囲み、「キー: 値」のペアを使います。例えば、`{ title: "Dune", year: 1965 }`のようにです。

4つの変数、`title`、`author`、`year`、`location`を宣言してください。それぞれに`trimmedParts`の対応する要素を代入します。次に、`return`の文を更新して、これら4つのプロパティを持つオブジェクトを返すようにしてください。

# --hints--

`parseCard`は、`title`、`author`、`year`、`location`のプロパティを持つオブジェクトを返す必要があります。

```js
const result = parseCard("The Stand | King, Stephen | 1978 | Shelf K2");
assert.isObject(result);
assert.property(result, "title");
assert.property(result, "author");
assert.property(result, "year");
assert.property(result, "location");
```

`title`が返すオブジェクトの`parseCard("The Stand | King, Stephen | 1978 | Shelf K2")`プロパティは、`"The Stand"`の値を持つ必要があります。

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").title, "The Stand");
```

`author`が返すオブジェクトの`parseCard("The Stand | King, Stephen | 1978 | Shelf K2")`プロパティは、`"King, Stephen"`の値を持つ必要があります。

```js
assert.strictEqual(parseCard("The Stand | King, Stephen | 1978 | Shelf K2").author, "King, Stephen");
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
--fcc-editable-region--

  return trimmedParts;
--fcc-editable-region--
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
