---
id: 699d1f8e51fcdc8230f779fb
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、遺産図書館のカタログシステムを作成します。生の書籍データを解析し、エントリーを検索・グループ化し、データの品質を検証し、結果をさまざまな形式でエクスポートします。

各カタログカードは、`"Title | Author | Year | Location"` 形式の単一の文字列として保存され、フィールドはパイプ文字（`|`）で区切られています。

このワークショップを通じて、これらの文字列の配列を使ってカタログ関数を構築およびテストします。`rawCatalogCards` という名前の配列を作成し、次の2つの文字列を追加してください。

```js
"From a Buick 8 | King, Stephen | 2002 | Shelf K7",
"The Shining | King, Stephen | 1977 | Shelf K1"
```

# --hints--

`rawCatalogCards` という名前の変数を宣言してください。

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+rawCatalogCards\s*=/);
```

`rawCatalogCards` は配列である必要があります。

```js
assert.isArray(rawCatalogCards);
```

`rawCatalogCards` の最初の要素は `"From a Buick 8 | King, Stephen | 2002 | Shelf K7"` であるべきです。

```js
assert.strictEqual(rawCatalogCards[0], "From a Buick 8 | King, Stephen | 2002 | Shelf K7");
```

`rawCatalogCards` の2番目の要素は `"The Shining | King, Stephen | 1977 | Shelf K1"` であるべきです。

```js
assert.strictEqual(rawCatalogCards[1], "The Shining | King, Stephen | 1977 | Shelf K1");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
