---
id: 699d20623e57e8ffee885459
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`parseCard` はテスト済みで正しく動作します。`cardResult` 変数はもう必要ありません。

まず、`const cardResult` 行と `console.log(cardResult)` 行を削除してください。それから、`parseCatalog` をパラメータに取る `rawCards` という名前の関数を作成してください。今のところ、関数内に空の配列 `catalog` を宣言し、それを返してください。

関数の下で `parseCatalog(rawCatalogCards)` を呼び出し、その結果を `catalog` という変数に代入し、`catalog.length` をコンソールに出力してください。

# --hints--

`rawCards` パラメータを持つ `parseCatalog` という名前の関数を作成してください。

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
assert.equal(parseCatalogFn.parameters[0]?.toString(), 'rawCards');
```

`catalog` という名前の変数を宣言し、空の配列で初期化してからそれを返してください。

```js
const explorer = await __helpers.Explorer(code);
const { parseCatalog: parseCatalogFn } = explorer.allFunctions;
assert.exists(parseCatalogFn);
const { catalog } = parseCatalogFn.variables;
assert.exists(catalog);
const catalogValue = eval(catalog.value?.toString());
assert.isArray(catalogValue);
assert.isEmpty(catalogValue);
assert.isTrue(parseCatalogFn.hasReturn('catalog'));
```

`catalog` は宣言され、`parseCatalog(rawCatalogCards)` に設定されているべきです。

```js
const explorer = await __helpers.Explorer(code);
const { catalog } = explorer.variables;
assert.exists(catalog);
assert.isTrue(catalog.value.matches('parseCatalog(rawCatalogCards)'));
```

`cardResult` 変数と関連する `console.log()` の呼び出しは削除してください。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bcardResult\b/);
```

`catalog.length` をコンソールにログ出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*catalog\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

--fcc-editable-region--
const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
--fcc-editable-region--
```
