---
id: 699d20623e57e8ffee885455
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

コンソールの出力からわかるように、分割した後でも各部分にはまだ余分な前後の空白があります。すべての要素をきれいにするには、それぞれを個別に処理する必要があります。

これは、配列のようなイテラブルの各要素に対してコードのブロックを1回ずつ繰り返す`for`ループを使うのに最適なケースです。

```js
for (let i = 0; i < array.length; i++) {
  // code to run on each iteration
}
```

`i`は`0`（最初のインデックス）から始まり、`i`が配列の長さより小さい間ループが続き、`i++`は各反復の後に`i`を1ずつ増やします。

`parseCard`の中で、`trimmedParts`という空の配列を宣言してください。次に`for`を繰り返す`parts`ループを追加します。ループの本体は今は空のままにしておいて、次のステップで埋めます。

# --hints--

`trimmedParts`という名前の変数を空の配列に設定して宣言してください。

```js
assert.match(parseCard.toString(), /(const|let|var)\s+trimmedParts\s*=\s*\[\s*\]/);
```

`for`を繰り返すために`parts`ループを使うべきです。

```js
assert.match(parseCard.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*parts\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
--fcc-editable-region--

--fcc-editable-region--
  return parts;
}

const cardResult = parseCard(rawCatalogCards[2]);
console.log(cardResult);
```
