---
id: 69b8f3a2c1d4e5f6a7b8c9d0
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

`parseCatalog` は現在空の配列を返しています。これを埋めるために、関数内に `for` を繰り返す `rawCards` ループを追加してください。ループの中で、各要素に対して `parseCard` を呼び出し、その結果を `catalog` にプッシュしてください。

# --hints--

`parseCatalog(rawCatalogCards).length` は `19` であるべきです。

```js
assert.strictEqual(parseCatalog(rawCatalogCards).length, 19);
```

`parseCatalog` は `for` を繰り返すために `rawCards` ループを使うべきです。

```js
assert.match(parseCatalog.toString(), /for\s*\(\s*(var|let)\s+(\w+)\s*=\s*0\s*;\s*\2\s*<\s*rawCards\s*\.\s*length\s*;\s*\2\+\+\s*\)\s*\{/);
```

ループの中で `parseCard` を呼び出すべきです。

```js
assert.match(parseCatalog.toString(), /parseCard\s*\(/);
```

# --seed--

## --seed-contents--

```js
const rawCatalogCards = [
  "From a Buick 8 | King, Stephen | 2002 | Shelf K7",
  "The Shining | King, Stephen | 1977 | Shelf K1",
  "The Stand | King, Stephen | 1978 | Shelf K2",
  "It | King, Stephen | 1986 | Shelf K3",
  "Misery | King, Stephen | 1987 | Shelf K4",
  "Do Androids Dream of Electric Sheep? | Dick, Philip K. | 1968 | Shelf D5",
  "I, Robot | Asimov, Isaac | 1950 | Shelf A8",
  "Foundation | Asimov, Isaac | 1951 | Shelf A9",
  "Dune | Herbert, Frank | 1965 | Shelf H3",
  "Neuromancer | Gibson, William | 1984 | Shelf G8",
  "Snow Crash | Stephenson, Neal | 1992 | Shelf S6",
  "The Martian | Weir, Andy | 2011 | Shelf W5",
  "Ender's Game | Card, Orson Scott | 1985 | Shelf C2",
  "The Hitchhiker's Guide to the Galaxy | Adams, Douglas | 1979 | Shelf A1",
  "Ready Player One | Cline, Ernest | 2011 | Shelf C7",
  "The Dark Tower: The Gunslinger | King, Stephen | 1982 | Shelf K5",
  // edge cases: missing data
  "Unknown Title |  | 1975 | Shelf X1",
  "Mysterious Manuscript | Unknown Author |  | Shelf Z9",
  "Ancient Scroll | Anonymous | 850 | ",
];

function parseCard(rawString) {
  const parts = rawString.split("|");
  const trimmedParts = [];
  for (let i = 0; i < parts.length; i++) {
    trimmedParts.push(parts[i].trim());
  }
  const title = trimmedParts[0];
  const author = trimmedParts[1];
  const year = trimmedParts[2];
  const location = trimmedParts[3];
  return {
    title: title || "Unknown",
    author: author || "Unknown",
    year: year ? parseInt(year) : "Unknown",
    location: location || "Unknown"
  };
}

function parseCatalog(rawCards) {
  const catalog = [];
--fcc-editable-region--

--fcc-editable-region--
  return catalog;
}

const catalog = parseCatalog(rawCatalogCards);
console.log(catalog.length);
```
