---
id: 66b62d0ad68488dd76228d6c
title: 偵錯隨機背景顏色變更器
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---

# --description--

Camperbot 正在學習 JavaScript，並嘗試建置他們自己的隨機背景顏色變更器。然而，他們在過程中犯了幾個錯誤。   

**目標：** 完成以下使用者故事並通過所有測試，使實驗室正常運作。

**使用者故事：**

1. 你應該修正 `darkColorsArr` 陣列中的語法錯誤。
1. 你應該修正當在 `getRandomIndex` 函式內取用 `math` 物件時所發出的 `ReferenceError`。
1. 你應該將 `randomIndex` 向下取整到最接近的整數。
1. 你應該更新 `body` 變數的值以使用正確的方法名稱並修正相關的 `TypeError`。
1. 你應該修正 `bgHexCodeSpanElement`，使其正確選取具有 `id` 為 `bg-hex-code` 的元素。
1. 你應該修改 `changeBackgroundColor` 中的 `color` 變數，使其設定為來自 `darkColorsArr` 陣列的隨機顏色。
1. 你應該更新 `btn` 變數，使其標的為按鈕元素。打開 `index.html` 以查看該按鈕元素的正確 `id` 名稱。
1. 你應該修正為 `btn` 添加 `click` 事件監聽器的那一行，將 `changeBackgroundColor` 函式的引用傳遞給該方法。
1. 每次點擊 `#btn` 時，你的背景應隨機變更為 `darkColorsArr` 中的其中一種顏色。

# --hints--

你應該解決 `darkColorsArr` 陣列中的語法錯誤。請記得陣列成員應該以逗號分隔。

```js
assert.isArray(darkColorsArr);
```

你應該修正 `math.random()` 行中的大小寫錯誤。 

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

你應該將 `darkColorsArr.length * Math.random()` 向下取整到最接近的整數。

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

你應該更新你的 `body` 變數的值，使 `queryselector` 使用駝峰式大小寫。

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

你應該修正 `"bg-hex-code"`，使其正確取用 `bg-hex-code` 的 `id`。請記得 CSS 中的 `id` 選擇器以井字號符號（`#`）開頭。

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

你應該在 `darkColorsArr[getRandomIndex]` 這一行中呼叫 `getRandomIndex` 函式（程式）。

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

你應該修正 `"#click-btn"` 行的 `id` 名稱，以符合 `index.html` 檔案中正確的 `id` 名稱。

```js
assert.match(btn?.id, /btn/);
```

你不應該呼叫 `changeBackgroundColor` 函式。相反地，你應該將函式引用傳遞給 `addEventListener` 方法。

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
