---
id: 6619240f46cec8e04d77e03a
title: Grundlegende Funktionen Übung A 
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

Erstellen Sie eine Funktion, die eine ganze Zahl entgegennimmt. Diese Funktion soll die gegebene `integer + 7` zurückgeben, wenn die ganze Zahl kleiner als `10` ist. Wenn die ganze Zahl größer oder gleich `10` ist, soll sie die gegebene `integer - 3` zurückgeben.

Der Name der Funktion soll `addOrSubtract` sein.

# --hints--

Sie sollten eine Funktion mit dem Namen `addOrSubtract` haben.

```js
assert.isFunction(addOrSubtract);
```

Ihre Funktion sollte eine ganze Zahl als Argument entgegennehmen.

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

Sie sollten die gegebene ganze Zahl + 7 zurückgeben, wenn die ganze Zahl kleiner als 10 ist.

```js
assert.strictEqual(addOrSubtract(5), 12);
```

Sie sollten die gegebene ganze Zahl - 3 zurückgeben, wenn die ganze Zahl größer oder gleich 10 ist.

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
