---
id: 661e17c6068359c3ccf2f4d8
title: Grundlegende Funktionen Übung D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

Schreiben Sie eine Funktion mit dem Namen `lastLetter`, die einen String als Parameter entgegennimmt und den letzten Buchstaben des Strings zurückgibt.

# --hints--

Sie sollten eine Funktion mit dem Namen `lastLetter` haben.

```js
assert.isFunction(lastLetter);
```

Ihre Funktion sollte einen String als Parameter entgegennehmen.

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

Sie sollten den letzten Buchstaben des Strings zurückgeben.

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
