---
id: 661e131f068359c3ccf2f4d6
title: Grundlegende Funktionen Übung B
challengeType: 1
dashedName: top-basic-functions-exercise-b
---

# --description--

Schreiben Sie eine Funktion mit dem Namen `multiply`, die zwei Parameter entgegennimmt und deren Produkt zurückgibt.

# --hints--

Sie sollten eine Funktion mit dem Namen `multiply` haben.

```js
assert.isFunction(multiply);
```

Ihre Funktion sollte zwei ganze Zahlen als Argumente entgegennehmen.

```js
assert.match(multiply.toString(), /\s*multiply\(\s*\w+\s*,\s*\w+\s*\)/);
```

Sie sollten das Produkt der beiden ganzen Zahlen zurückgeben.

```js
assert.strictEqual(multiply(10, 10), 100);
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js 
function multiply(a, b) {
  return a * b;
}
```
