---
id: 661e151f068359c3ccf2f4d7
title: Grundlegende Funktionen Übung C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

Schreiben Sie eine Funktion mit dem Namen `capitalize`, die einen String als Parameter entgegennimmt und einen neuen String mit dem großgeschriebenen ersten Buchstaben zurückgibt.

# --hints--

Sie sollten eine Funktion mit dem Namen `capitalize` haben.

```js
assert.isFunction(capitalize);
```

Ihre Funktion sollte einen String als Parameter entgegennehmen.

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

Ihre Funktion sollte einen neuen String mit dem großgeschriebenen ersten Buchstaben zurückgeben.

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
