---
id: 5987fd532b954e0f21b5d3f6
title: Equilibrium index
challengeType: 1
forumTopicId: 302255
dashedName: equilibrium-index
---

# --description--

Um índice de equilíbrio de uma sequência é um índice na sequência, tal que a soma dos elementos nos índices mais baixos é igual à soma dos elementos nos índices mais altos.

Por exemplo, em uma sequência $A$:

<ul style='list-style: none;'>
  <li><big>$A_0 = -7$</big></li>
  <li><big>$A_1 =  1$</big></li>
  <li><big>$A_2 =  5$</big></li>
  <li><big>$A_3 =  2$</big></li>
  <li><big>$A_4 = -4$</big></li>
  <li><big>$A_5 =  3$</big></li>
  <li><big>$A_6 =  0$</big></li>
</ul>

`3` é um índice de equilíbrio, porque:

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 = A_4 + A_5 + A_6$</big></li>
</ul>

`6` também é um índice de equilíbrio, porque:

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 + A_3 + A_4 + A_5 = 0$</big></li>
</ul>

(a soma de zero elementos é zero)

`7` não é um índice de equilíbrio porque não é um índice válido da sequência $A$.

# --instructions--

Escreva uma função que, dada uma sequência, retorne seus índices de equilíbrio (se houver).

Suponha que a sequência seja muito longa.

# --hints--

`equilibrium` deve ser uma função.

```js
assert(typeof equilibrium === 'function');
```

`equilibrium([-7, 1, 5, 2, -4, 3, 0])` should return `[3,6]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[0]),
  globalThis.ans[0]
);
```

`equilibrium([2, 4, 6])` should return `[]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[1]),
  globalThis.ans[1]
);
```

`equilibrium([2, 9, 2])` should return `[1]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[2]),
  globalThis.ans[2]
);
```

`equilibrium([1, -1, 1, -1, 1, -1, 1])` should return `[0,1,2,3,4,5,6]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[3]),
  globalThis.ans[3]
);
```

`equilibrium([1])` should return `[0]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[4]),
  globalThis.ans[4]
);
```

`equilibrium([])` should return `[]`.

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[5]),
  globalThis.ans[5]
);
```

# --before-all--

```js
globalThis.equilibriumTests =
  [[-7, 1, 5, 2, -4, 3, 0], // 3, 6
  [2, 4, 6], // empty
  [2, 9, 2], // 1
  [1, -1, 1, -1, 1, -1, 1], // 0,1,2,3,4,5,6
  [1], // 0
  [] // empty
  ];
globalThis.ans = [[3, 6], [], [1], [0, 1, 2, 3, 4, 5, 6], [0], []];
```

# --seed--

## --seed-contents--

```js
function equilibrium(a) {

}
```

# --solutions--

```js
function equilibrium(a) {
  let N = a.length,
    i,
    l = [],
    r = [],
    e = [];
  for (l[0] = a[0], r[N - 1] = a[N - 1], i = 1; i < N; i++)
    { l[i] = l[i - 1] + a[i], r[N - i - 1] = r[N - i] + a[N - i - 1]; }
  for (i = 0; i < N; i++)
    { if (l[i] === r[i]) e.push(i); }
  return e;
}
```
