---
id: 5a23c84252665b21eecc7e06
title: Cortar um retângulo
challengeType: 1
forumTopicId: 302242
dashedName: cut-a-rectangle
---

# --description--

Um retângulo dado é feito de *m* × *n* quadrados. Se *m* e *n* não forem ambos ímpares, então é possível cortar um caminho através do retângulo ao longo das bordas dos quadrados de forma que o retângulo se divida em duas peças conectadas com a mesma forma (após girar uma das peças em 180°). Todos os caminhos para retângulos 2 × 2 e 4 × 3 são mostrados abaixo.

<img src="https://cdn.freecodecamp.org/curriculum/coding-interview-prep/rosetta-cut-a-rectangle.png" alt="Diagram showing the possible paths for 2 by 2 and 4 by 3 rectangles" />

# --instructions--

Escreva uma função que calcule o número de maneiras diferentes de cortar um retângulo *m* × *n*.

# --hints--

`cutRectangle` deve ser uma função.

```js
assert(typeof cutRectangle == 'function');
```

`cutRectangle(2, 2)` deve retornar um número.

```js
assert(typeof cutRectangle(2, 2) == 'number');
```

`cutRectangle(2, 2)` deve retornar `2`.

```js
assert.equal(cutRectangle(2, 2), 2);
```

`cutRectangle(4, 3)` deve retornar `9`.

```js
assert.equal(cutRectangle(4, 3), 9);
```

`cutRectangle(4, 4)` deve retornar `22`.

```js
assert.equal(cutRectangle(4, 4), 22);
```

`cutRectangle(8, 3)` deve retornar `53`.

```js
assert.equal(cutRectangle(8, 3), 53);
```

`cutRectangle(7, 4)` deve retornar `151`.

```js
assert.equal(cutRectangle(7, 4), 151);
```

# --seed--

## --seed-contents--

```js
function cutRectangle(w, h) {

}
```

# --solutions--

```js
function cutRectangle(w, h) {
  if (w % 2 == 1 && h % 2 == 1) return;

  var dirs = [[0, -1], [-1, 0], [0, 1], [1, 0]];

  var grid = new Array(h);
  for (var i = 0; i < grid.length; i++) grid[i] = new Array(w);
  var stack = [];

  var half = Math.floor((w * h) / 2);
  var bits = Math.pow(2, half) - 1;
  var result = 0;
  for (; bits > 0; bits -= 2) {
    for (var i = 0; i < half; i++) {
      var r = Math.floor(i / w);
      var c = i % w;
      grid[r][c] = (bits & (1 << i)) != 0 ? 1 : 0;
      grid[h - r - 1][w - c - 1] = 1 - grid[r][c];
    }

    stack.push(0);
    grid[0][0] = 2;
    var count = 1;
    while (stack.length != 0) {
      var pos = stack.pop();
      var r = Math.floor(pos / w);
      var c = pos % w;

      for (var dir of dirs) {
        var nextR = r + dir[0];
        var nextC = c + dir[1];

        if (nextR >= 0 && nextR < h && nextC >= 0 && nextC < w) {
          if (grid[nextR][nextC] == 1) {
            stack.push(nextR * w + nextC);
            grid[nextR][nextC] = 2;
            count++;
          }
        }
      }
    }

    if (count == half) {
      result++;
    }
  }

  return result;
}
```
