---
id: 5e6dd14797f5ce267c2f19d0
title: Sequência para olhar e dizer
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

A sequência para olhar e dizer é uma sequência de números recursivamente definida.

Definição da sequência

<ul><li>Take a decimal number</li>
<li><span>Look</span> at the number, visually grouping consecutive runs of the same digit.</li>
<li><span>Say</span> the number, from left to right, group by group; as how many of that digit there are - followed by the digit grouped.</li></ul><span> This becomes the next number of the sequence.</span>

Um exemplo:

<ul><li>Starting with the number 1, you have <span>one</span> 1 which produces 11</li>
<li>Starting with 11, you have <span>two</span> 1's. I.E.: 21</li>
<li>Starting with 21, you have <span>one</span> 2, then <span>one</span> 1. I.E.: (12)(11) which becomes 1211</li>
<li>Starting with 1211, you have <span>one</span> 1, <span>one</span> 2, then <span>two</span> 1's. I.E.: (11)(12)(21) which becomes 111221</li></ul>

# --instructions--

Escreva uma função que aceita uma string como parâmetro, faz seu processamento e retorna a string resultante.

# --hints--

`lookAndSay` deve ser uma função.

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` deve retornar uma string.

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` deve retornar `"11"`.

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` deve retornar `"21"`.

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` deve retornar `"1211"`.

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` deve retornar `"111221"`.

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` deve retornar `"13151412"`.

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
