---
id: 66c63bfa4685e4e3843fa0dc
title: 步骤 1
challengeType: 1
dashedName: step-1
---

# --description--

在本工作坊中，你将通过创建一个杂货购物清单来继续学习数组。

首先添加一个 `console.log` 在控制台记录字符串 `"Grocery shopping list"`。

# --hints--

你的代码中应该有一个 `console.log`。

```js
assert.match(code, /console\.log(.*)/);
```

你的 `console.log` 应输出信息 `"Grocery shopping list"`。 仔细检查拼写和大小写错误。

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
