---
id: 66c64ad16796c7f2419b45c5
title: 步骤 8
challengeType: 1
dashedName: step-8
---

# --description--

现在是时候开始在杂货清单前面添加物品了。

首先添加一个 `console.log`，将信息 `"It looks like we need to get some cooking oil."` 记录到控制台。

# --hints--

你应该在控制台中记录消息 `"It looks like we need to get some cooking oil."`。

```js
assert.match(code, /console\.log\((['"])(It\s+looks\s+like\s+we\s+need\s+to\s+get\s+some\s+cooking\s+oil\.)\1\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
