---
id: 66c6491fe4c8e0f16845425f
title: 步骤7
challengeType: 1
dashedName: step-7
---

# --description--

现在，是时候把另一种水果列入清单了。

使用与之前相同的数组方法，将字符串 `"Grapes"` 添加到 `shoppingList` 数组的末尾。

然后，添加 `console.log`，并在 `console.log` 中调用 `getShoppingListMsg` 函数，查看控制台记录的更新列表。

# --hints--

你的 `shoppingList` 数组现在应包含两个项目。

```js
assert.lengthOf(shoppingList, 2);
```

你的 `shoppingList` 数组应包字符串 `"Grapes"`。

```js
assert.equal(shoppingList[1], "Grapes");
```

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)` 函数。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
