---
id: 66c644c08ebcd2ef90c750bd
title: 步骤4
challengeType: 1
dashedName: step-4
---

# --description--

在之前的课程中，你学习了如何使用 `push` 方法将项添加到数组的末尾，方法如下：

```js
fruits.push('pear');
```

使用 `push` 方法，将字符串 `"Apples"` 添加到 `shoppingList` 数组。

# --hints--

你的 `shoppingList` 数组中应该有一个项目。

```js
assert.lengthOf(shoppingList, 1);
```

你的 `shoppingList` 数组中应包含字符串 `"Apples"`。 仔细检查拼写和大小写错误。

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
