---
id: 66c64095f61166ede6e39a84
title: 步骤2
challengeType: 1
dashedName: step-2
---

# --description--

在这份购物清单中，你将使用一个数组来表示你需要购买的物品。

在之前的课程中，你学习了如何像这样创建数组：

```js
const fruits = ['apple', 'banana', 'orange'];
```

在此步骤中，创建一个名为 `shoppingList` 的变量，并为其赋值一个空数组。

# --hints--

你应该有一个名为 `shoppingList` 的变量。

```js
assert.isNotNull(shoppingList);
```

你应该为 `shoppingList` 变量赋值一个数组。

```js
assert.isArray(shoppingList);
```

你的数组应该是空的。

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
