---
id: 66c645b345a39bf04864dc50
title: 步骤5
challengeType: 1
dashedName: step-5
---

# --description--

要查看更新后的购物清单，你需要将当前购物清单和一条简短信息记录到控制台。

由于这条信息将在整个工作坊中反复使用，因此最好创建一个可重复使用的函数。

创建一个名为 `getShoppingListMsg` 的函数，将数组作为参数并返回字符串 `"Current Shopping List: "` 和所提供数组的内容。

你可以使用模板字面量或使用 `+` 操作符进行字符串连接，将字符串与传递给函数的数组结合起来。

# --hints--

你的代码中应该有一个名为 `getShoppingListMsg` 的函数。

```js
assert.isFunction(getShoppingListMsg);
```

你的 `getShoppingListMsg` 函数应返回一个字符串。

```js
assert.isString(getShoppingListMsg(shoppingList));
```

你的 `getShoppingListMsg` 函数应接受一个数组作为参数。


```js
assert.strictEqual(getShoppingListMsg(["Anchovies", "Olives"]), "Current Shopping List: Anchovies,Olives");
```

你的 `getShoppingListMsg` 函数应接受一个数组并返回信息 `"Current Shopping List: "` 和该数组的内容。 仔细检查空格和拼写错误。

```js
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
