---
id: 66c64210413532ee9d3bd342
title: 步骤 3
challengeType: 1
dashedName: step-3
---

# --description--

接下来，你将练习在杂货清单中添加食品。

首先使用 `console.log` 记录信息 `"It will be nice to have some fruit to eat."`。 

# --hints--

你应将信息 `"It will be nice to have some fruit to eat."` 记录到控制台。

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
