---
id: 66c73a7798f6f62b2ae58f22
title: 步骤 16
challengeType: 1
dashedName: step-16
---

# --description--

现在是时候在杂货清单的开头添加更多物品了。

首先添加 `console.log` 语句，记录信息 `"It might be nice to get a dessert."`

在 `console` 语句下方，使用正确的数组方法将字符串 `"Chocolate Cake"` 添加到 `shoppingList` 数组的开头。

最后，添加 `console.log`，并在 `console.log` 中调用 `getShoppingListMsg` 函数，查看控制台记录的更新列表。

# --hints--

你应该有一个 `console.log` 来记录信息 `"It might be nice to get a dessert."`。

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

你的 `shoppingList` 数组总共应有 `6` 个项目。

```js
assert.lengthOf(shoppingList, 6);
```

你应该使用 `unshift` 方法将字符串 `"Chocolate Cake"` 添加到 `shoppingList` 数组的开头。

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
