---
id: 66c726c34ecf1e238aa9d7d5
title: 步骤 9
challengeType: 1
dashedName: step-9
---

# --description--

在之前的课程中，你学习了如何使用 `unshift()` 方法将元素添加到数组的开头。 

下面是提醒你如何使用 `unshift()` 方法：

```js
array.unshift(item1, item2, ..., itemX);
```

使用 `unshift()` 方法将字符串 `"Vegetable Oil"` 添加到 `shoppingList` 数组的开头。

# --hints--

你的 `shoppingList` 数组中应该有三个项目。

```js
assert.lengthOf(shoppingList, 3);
```

你应该使用 `unshift()` 方法将 `"Vegetable Oil"` 添加到 `shoppingList` 数组的开头。

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
