---
id: 66c730ee6ae076281721d0b9
title: 步骤14
challengeType: 1
dashedName: step-14
---

# --description--

In previous lessons, you learned how to remove items from the end of an array using the `pop` method.

下面是提醒你如何使用 `pop` 方法：

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

使用 `pop` 方法从 `shoppingList` 数组删除最后一个项目。

# --hints--

你的 `shoppingList` 数组现在应该有 `5` 项目。

```js
assert.lengthOf(shoppingList, 5);
```

你的 `shoppingList` 数组中不应再包含字符串 `"Potato Chips"`。

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
