---
id: 66c72a55418cc9247b710827
title: 步骤 10
challengeType: 1
dashedName: step-10
---

# --description--

接下来，添加 `console.log`，并在 `console.log` 中调用 `getShoppingListMsg` 函数，查看控制台记录的更新列表。

# --hints--

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)`。

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 3);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");

--fcc-editable-region--

--fcc-editable-region--
```
