---
id: 66c74079c30b1c2d166cb9a4
title: 步骤18
challengeType: 1
dashedName: step-18
---

# --description--

在之前的课程中，你学习了如何使用 `shift` 方法从数组开头移除一项。 

下面是提醒你如何使用 `shift` 方法：

```js
const array = [1, 2, 3, 4, 5];
array.shift();

// Result: [2, 3, 4, 5]
console.log(array); 
```

使用 `shift` 方法从 `shoppingList` 数组删除第一个项目。

# --hints--

你的 `shoppingList` 数组总共应有 `5` 个项目。

```js
assert.lengthOf(shoppingList, 5);
```

你应使用 `shift` 方法从 `shoppingList` 数组中移除第一个项目。

```js
assert.strictEqual(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

--fcc-editable-region--

--fcc-editable-region--
```
