---
id: 66cbe2319d3845545a293a0b
title: 步骤6
challengeType: 1
dashedName: step-6
---

# --description--

现在可以查看记录到控制台的信息了。

添加一个 `console.log`，并在 `console.log` 中调用 `getShoppingListMsg` 函数，查看记录到控制台的信息。

# --hints--

你应该在 `console.log` 中调用 `getShoppingListMsg(shoppingList)` 函数。

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
