---
id: a789b3483989747d63b0e427
title: Costruisci il cercatore del numero più grande
challengeType: 26
dashedName: build-the-largest-number-finder
---

# --description--

In questo laboratorio costruirai una funzione che restituisce un array composto dal numero più grande di ogni sotto-array fornito.

Ricorda, puoi iterare su un array con un semplice loop `for` e consultare ogni elemento con la sintassi degli array `arr[i]`.

**Obiettivo:** Soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una funzione `largestOfAll` che prende come argomento un array di array.
2. La funzione dovrebbe restituire un array contenente il numero più grande di ogni sotto-array.

# --hints--

Dovresti creare una funzione chiamata `largestOfAll`.

```js
assert.isFunction(largestOfAll);
```

`largestOfAll` dovrebbe prendere un singolo parametro.

```js
assert.lengthOf(largestOfAll, 1);
```

`largestOfAll([[4, 5, 1, 3], [13, 27, 18, 26], [32, 35, 37, 39], [1000, 1001, 857, 1]])` dovrebbe restituire un array.

```js
assert.isArray(
  largestOfAll([
    [4, 5, 1, 3],
    [13, 27, 18, 26],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ])
);
```

`largestOfAll([[13, 27, 18, 26], [4, 5, 1, 3], [32, 35, 37, 39], [1000, 1001, 857, 1]])` dovrebbe restituire `[27, 5, 39, 1001]`.

```js
assert.deepEqual(
  largestOfAll([
    [13, 27, 18, 26],
    [4, 5, 1, 3],
    [32, 35, 37, 39],
    [1000, 1001, 857, 1]
  ]),
  [27, 5, 39, 1001]
);
```

`largestOfAll([[4, 9, 1, 3], [13, 35, 18, 26], [32, 35, 97, 39], [1000000, 1001, 857, 1]])` dovrebbe restituire `[9, 35, 97, 1000000]`.

```js
assert.deepEqual(
  largestOfAll([
    [4, 9, 1, 3],
    [13, 35, 18, 26],
    [32, 35, 97, 39],
    [1000000, 1001, 857, 1]
  ]),
  [9, 35, 97, 1000000]
);
```

`largestOfAll([[17, 23, 25, 12], [25, 7, 34, 48], [4, -10, 18, 21], [-72, -3, -17, -10]])` dovrebbe restituire `[25, 48, 21, -3]`.

```js
assert.deepEqual(
  largestOfAll([
    [17, 23, 25, 12],
    [25, 7, 34, 48],
    [4, -10, 18, 21],
    [-72, -3, -17, -10]
  ]),
  [25, 48, 21, -3]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function largestOfAll(arr) {
  return arr.map(subArr => Math.max.apply(null, subArr));
}

largestOfAll([
  [4, 5, 1, 3],
  [13, 27, 18, 26],
  [32, 35, 37, 39],
  [1000, 1001, 857, 1]
]);
```
