---
id: 68dbc6b7e8f794a2ae69bc74
title: N-퀸 알고리즘 구현하기
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

N-퀸 문제는 N×N 체스판 위에 N개의 퀸을 서로 공격하지 못하도록(같은 행, 열, 대각선에 두지 않도록) 배치하는 문제입니다.

예를 들어, 4x4 체스판이 있을 때, 유효한 배치 중 하나는 다음과 같습니다:

```md
[1, 3, 0, 2]
```

이는 0행에는 퀸이 1열에, 1행에는 3열에, 2행에는 0열에, 3행에는 2열에 배치되었음을 의미합니다.

시각적으로 이 배치는 다음과 같습니다:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

여기서 `Q`는 퀸을, `.`는 빈 칸을 나타냅니다.

이 실습에서는 깊이 우선 탐색 방식을 사용해 N-퀸 문제 해결기를 구현할 것입니다.


**목표**: 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `dfs_n_queens`라는 이름의 함수가 있어야 합니다.
2. 함수는 정확히 하나의 인수, 정수 `n`를 받아야 합니다.
3. 만약 `n`가 `1`보다 작으면, 함수는 빈 리스트(`[]`)를 반환해야 합니다.
4. 함수는 해답들의 리스트를 반환해야 하며, 각 해답은 길이가 `n`인 리스트이고, 인덱스 `i`의 요소는 `i` 행에 있는 퀸의 0부터 시작하는 열 인덱스입니다.

# --hints--

하나의 인수를 받는 `dfs_n_queens`라는 함수가 있어야 합니다.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

만약 `n`이 `1`보다 작으면, 함수는 빈 리스트를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

함수는 길이가 `n`인 리스트들로 이루어진 해답 리스트를 반환해야 합니다.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)`는 `[[0]]`를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)`는 `[]`를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)`는 `[]`를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)`는 `[[1, 3, 0, 2], [2, 0, 3, 1]]`를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)`는 `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))`는 `10`여야 합니다.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))`는 `92`여야 합니다.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens`는 올바른 결과를 반환해야 합니다.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
