---
id: 619665c9abd72906f3ad30f9
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Estarás construyendo un feliz pingüino Flappy y explorando aún más las transformaciones y animaciones CSS en el proceso.

Comience con su plantilla HTML básica. Incluya la declaración `DOCTYPE`, el elemento `html` con un idioma establecido en inglés, las etiquetas `meta` apropiadas, un `head`, `body` y elemento `title`. Además, vincule su hoja de estilo a la página.

# --hints--

Tu código debe tener una declaración `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

El código debe tener un elemento `html`.

```js
assert.equal(document.querySelectorAll('html')?.length, 1);
```

Tu elemento `html` debe tener una etiqueta de apertura con el atributo `lang` con el valor `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Su elemento `html` debe tener una etiqueta de cierre.

```js
assert(code.match(/<\/html\s*>/gi));
```

Tu código debe tener un elemento `head` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('head')?.length, 1);
```

Tu código debe tener un elemento `body` dentro del elemento `html`.

```js
assert.equal(document.querySelectorAll('body')?.length, 1);
```

Tu elemento `head` debe estar antes de tu elemento `body`.

```js
assert.equal(document.querySelector('body')?.previousElementSibling?.tagName, 'HEAD');
```

Debes tener dos elementos `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.equal(meta?.length, 2);
```

Un elemento `meta` debe tener un `name` establecido en `viewport`, y el `content` establecido a `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

El otro elemento `meta` debe tener el atributo `charset` establecido en `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu `title` debe tener un poco de texto.

```js
const title = document.querySelector('title');
assert.isAtLeast(title?.textContent?.length, 1);
```

Su código debe tener un elemento <code>link`.

```js
assert.exists(document.querySelector('link'));
```

Su elemento `link` debe estar dentro de su elemento `head`.

```js
assert(code.match(/<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i));
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
