---
id: 66629f407d679d3105e8317f
title: 构建一个猜拳游戏
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
你的游戏将在电脑上运行。 你应该编写一个可以随机返回 `"rock"` 、`"paper"` 或 `"scissors"` 的函数。

你不需要担心游戏的前端部分。 你只需要编写游戏的逻辑部分。 打开 `script.js` 并开始按照指示操作。

**用户故事：**

1. 你应该有一个名为 `getComputerChoice` 的函数。

1. `getComputerChoice` 函数应该随机返回 `"rock"` 、`"paper"` 或 `"scissors"` 。

**提示：** `[Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random)` 方法返回一个大于或相等于 0 且小于 1 的随机数。思考如何使用它来条件式地返回多个选项中的一个。

你的游戏将由一个用户游玩。 你将要编写一个可以获取用户选择并返回的函数。

1. 创建一个名为 `getHumanChoice` 的函数。

1. 编写 `getHumanChoice` 函数的代码，使函数可以返回用户输入的有效选择。

**提示：** 使用 [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) 方法获取用户的输入。

你的游戏将对用户得分进行记录。 你应该编写变量来记录玩家的得分。

1. 创建两个名为 `humanScore` 和 `computerScore` 的全局变量。

1. 将变量的值初始化为 `0` 。

你的游戏是回合制的。 你将编写一个将用户和电脑选择作为参数，可以进行一轮游戏、增加每一轮获胜者的分数、记录获胜者公告的函数。

1. 创建一个名为 `playRound` 的函数。

1. 为 `playRound` 函数定义两个参数。 `humanChoice` 参数和 `computerChoice` 参数。 通过这两个参数来获取用户和电脑的选择。

1. 使你的函数 `humanChoice` 参数可以忽略大小写，这样玩家就可以输入 `"rock"` 、`"ROCK"` 、 `"RocK"` 或者其他输入。

1. 编写 `playRound` 函数的代码，返回代表获胜者的字符串。

- 如果平局，应该返回 `"It's a tie!"`。

- 如果玩家获胜，它应该返回 `"You win! [player choice] beats [computer choice]"`。

- 如果计算机获胜，它应该返回 `"You lose! [computer choice] beats [player choice]"`。

1. 根据回合获胜者，增加 `humanScore` 变量或者 `computerScore` 变量。

游戏将进行三轮。 你应该编写一个 `playGame` 函数，该函数应该调用三次 `playRound` 、记录得分和宣布获胜者。

1. 创建一个名为 `playGame` 的函数。

1. 创建一个循环三次的循环，并且在循环中调用以用户选择和电脑选择作为参数的 `playRound` 函数。

1. 在游戏结束时，根据谁赢得了最多的回合，返回游戏的获胜者。

- 如果用户比电脑获胜的回合多，返回 `"You win the game!"` 消息。

- 如果电脑比用户获胜的回合多，返回 `"You lose the game!"` 消息。

# --hints--

你应该有一个名为 `getComputerChoice` 的函数

```js
assert.isFunction(getComputerChoice);
```

`getComputerChoice` 函数应该随机返回 `"rock"` 、`"paper"` 或 `"scissors"` 。

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

你应该有一个名为 `getHumanChoice` 的函数

```js
assert.isFunction(getHumanChoice);
```

你应该有两个名为 `humanScore` 和 `computerScore` 的全局变量。

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

你应该有一个名为 `playRound` 的函数

```js
assert.isFunction(playRound);
```

你的 `playRound` 函数应以参数 `humanChoice` 和 `computerChoice` 作为参数。

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

你的 `playRound` 函数参数应该可以忽略大小写，这样玩家就可以输入 `"rock"` 、`"ROCK"` 、 `"RocK"` 或者其他输入。

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

如果用户和电脑的选择相同， `playRound` 函数应该返回 `"It's a tie!"` 。

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

你的 `playRound` 函数如果玩家获胜，应返回字串 `"You win! [player choice] beats [computer choice]"`。

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

你的 `playRound` 函数如果计算机获胜，应返回字串 `"You lose! [computer choice] beats [player choice]"`。

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

`playRound` 函数应该根据每轮的获胜者增加 `humanScore` 或 `computerScore` 的得分。

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

你应该有一个名为 `playGame` 的函数。

```js
assert.isFunction(playGame);
```

你应该循环3次。

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

你应该根据获胜的轮数返回游戏的获胜者。

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```

