---
id: 66f686b8ebdb982fa8e14330
title: 建置一個日期轉換程式
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

在這個實驗中，你將撰寫一個日期轉換程式，將給定的日期轉換成不同的格式。 例如，當前日期 `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` 將會以以下 2 種方式格式化：

- (MM/DD/YYYY)：9/27/2024。
- （年 月 日）：2024 年 9 月 27 日。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該創建一個名為 `currentDate` 的變數，並使用 `Date` 物件將當前日期和時間指定給它。
2. 你應該創建一個名為 `currentDateFormat` 的變數，並將字串 `Current Date and Time: [current date]` 指定給它。將 `[current date]` 替換為 `currentDate` 變數的結果。
3. 你應該將 `currentDateFormat` 的值記錄到主控台。
4. 你應該創建一個名為 `formatDateMMDDYYYY` 的函式（程式），該函式（程式）以日期物件作為參數。你可以將此參數命名為任何你喜歡的名稱。
5. 你的 `formatDateMMDDYYYY` 函式（程式）應該將目前日期轉換為 `month/day/year` 格式，並傳回字串 `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`。
6. 你應該創建一個名為 `formatDateLong` 的函式（程式），該函式（程式）以日期物件作為參數。
7. 你的 `formatDateLong` 函式（程式）應該將目前日期轉換為 `Month Day, Year` 格式，並傳回字串 `Formatted Date (Month Day, Year): [formatted date]`。

**注意：** 為了讓測試通過，請確保在格式化日期時使用 `en-US` 作為區域的設定。

# --hints--

你應該有一個名為 `currentDate` 的變數，該變數使用 `Date` 物件來保存目前的日期和時間。

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

你應該有一個名為 `currentDateFormat` 的變數，用來保存格式為 `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>` 的當前日期。

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

你應該將 `currentDateFormat` 的值記錄到主控台。

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

你應該有一個函式（程式）`formatDateMMDDYYYY` 

```js
assert.isFunction(formatDateMMDDYYYY);
```

函式（程式）`formatDateMMDDYYYY` 應該接受一個參數。

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

當日期物件持有 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 時，函式（程式）`formatDateMMDDYYYY` 應該傳回 `Formatted Date (MM/DD/YYYY): 9/27/2024`。

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

你應該有一個函式（程式）`formatDateLong` 

```js
assert.isFunction(formatDateLong);
```

函式（程式）`formatDateLong` 應該接受一個單一的參數。

```js
assert.lengthOf(formatDateLong, 1);
```

當日期物件持有 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)` 時，函式（程式）`formatDateLong` 應該傳回 `Formatted Date (Month Day, Year): September 27, 2024`。

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

