---
id: 694c84330cbe604821b9deef
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kituo cha mawasiliano kinahitaji njia ya kuongeza wanachama wapya kwenye orodha. Unda kitendakazi tupu kinachoitwa `addCrewMember` kinachokubali vigezo viwili, `crew` na `astronaut`.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `addCrewMember`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addCrewMember);
```

Kitendakazi chako cha `addCrewMember` kinapaswa kuwa na vigezo `crew` na `astronaut`.

```js
const explorer = await __helpers.Explorer(code);
const { addCrewMember } = explorer.allFunctions;
assert.lengthOf(addCrewMember?.parameters, 2);
assert.equal(addCrewMember?.parameters?.[0]?.toString(), "crew");
assert.equal(addCrewMember?.parameters?.[1]?.toString(), "astronaut");
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
