---
id: 694c84106e9023f289c03b84
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Orodha yako inaanza na mwanamkubwa mmoja tu wa wanadamu wa anga. Unda kitu kinachoitwa `firstAstronaut` chenye sifa zifuatazo:

| Key | Value |
| ----------- | ------- |
|`id`|`1`|
|`name`|`"Andy"`|
|`role`|`"Commander"`|
|`isEVAEligible`|`true`|
|`priority`|`3`|

# --hints--

Unapaswa kuunda tofauti ya `firstAstronaut`.

```js
assert.isDefined(firstAstronaut);
```

`firstAstronaut` yako inapaswa kuwa kitu.

```js
assert.isObject(firstAstronaut);
```

Kitu chako cha `firstAstronaut` kinapaswa kujumuisha thamani sahihi za sifa.

```js
assert.deepEqual(firstAstronaut, { id: 1, name: "Andy", role: "Commander", isEVAEligible: true, priority: 3 });
```

# --seed--

## --seed-contents--

```js
const squad = [];

--fcc-editable-region--

--fcc-editable-region--
```
