---
id: 694c40d86fe6a78d20b6559c
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utatumia JavaScript kujenga orodha ya wanaanga kwa ajili ya misheni ya anga. 

Kuanza, tengeneza safu ya data tupu kwa ajili ya orodha yako iitwayo `squad`.

# --hints--

Unapaswa kutengeneza variable ya `squad`.

```js
assert.isDefined(squad);
```

Variable yako ya `squad` inapaswa kuwa safu ya data.

```js
assert.isArray(squad);
```

Safu yako ya `squad` inapaswa kuwa tupu.

```js
assert.deepEqual(squad, []);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
