---
id: 694c84330cbe604821b9deef
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Kituo cha mawasiliano kinahitaji njia ya kuongeza wanachama wapya kwenye orodha. Unda kitendakazi tupu kinachoitwa `addCrewMember` kinachokubali vigezo viwili, `crew` na `astronaut`.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `addCrewMember`.

```js
assert.isFunction(addCrewMember);
```

Kitendakazi chako cha `addCrewMember` kinapaswa kuwa na vigezo `crew` na `astronaut`.

```js
const regex = __helpers.functionRegex('addCrewMember', ['crew', 'astronaut']);
assert.match(__helpers.removeJSComments(addCrewMember.toString()), regex);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

--fcc-editable-region--

--fcc-editable-region--
```
