---
id: 694d8ddd106340237efb0c80
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Njia ya `splice()` inaweza kubadilisha safu za data kwa kuongeza au kuondoa vipengele mahali popote, ikiwa ni pamoja na katikati. Kwa kuwa `splice()` inarudisha safu ya data yenye vipengele vilivyotolewa, unaweza kuitumia kubadilisha vipengele viwili katika safu ya data bila kubadilisha asili yake. Hii inaweza kufanywa kwa mstari mmoja kwa kutumia mbinu ifuatayo:

```js
// swap elements at i and j without mutating the original
const copy = array.slice();
copy[i] = copy.splice(j, 1, copy[i])[0];
```

Hapa kuna mfano:

```js
const originalArray = [12, 97, 68, 55];
const copyArray = originalArray.slice();
copyArray[1] = copyArray.splice(3, 1, copyArray[1])[0];
console.log(copyArray); // [12, 55, 68, 97]
```

Mbinu, inayotumika katika mstari wa tatu, hufanya kazi kama ifuatavyo:

- `splice(3, 1, copyArray[1])` hutoa kipengele kilichoko kwenye kielezo `3` (`55`)

- Inaingiza kipengele kutoka `copyArray[1]` (`97`) kwenye kielezo `3`

- `splice()` inarudisha `[55]`, safu ya data yenye kipengele kilichotolewa

- `[0]` hutoa `55` kutoka kwenye safu hiyo

- Thamani hiyo (`55`) imepewa tena `copyArray[1]`, ikikamilisha kubadilisha

Tumia mbinu hii na safu ya `updatedCrew` kubadilisha nafasi za wanahabari wa anga walioko `fromIndex` na `toIndex`.

# --hints--

Unapaswa kufanya kubadilisha kwa kuita njia ya `splice()`.

```js
assert.match(swapCrewMembers.toString(), /splice\s*\(/);
```

Unapaswa kuita `splice()` na safu ya `updatedCrew` kufanya kubadilisha. Unapaswa kutumia mbinu ya mstari mmoja.

```js
const cleaned = __helpers.removeWhiteSpace(__helpers.removeJSComments(swapCrewMembers.toString()));
const regex = /updatedCrew\[(fromIndex|toIndex)\]=updatedCrew\.splice\((fromIndex|toIndex),1,updatedCrew\[\1\]\)\[0\];?\s*\}$/;
assert.match(cleaned, regex);
```

Kitendakazi chako cha `swapCrewMembers` hakipaswi kubadilisha safu ya data ya ingizo `crew` asili.

```js
const demo = [{ id: 1 }, { id: 2 }, { id: 3 }];
swapCrewMembers(demo, 0, 2);
assert.deepEqual(demo.map(a => a.id), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
const squad = [];

const firstAstronaut = {
  id: 1,
  name: "Andy",
  role: "Commander",
  isEVAEligible: true,
  priority: 3
};

function addCrewMember(crew, astronaut) {
  for (let i = 0; i < crew.length; i++) {
    if (crew[i].id === astronaut.id) {
      console.log("Duplicate ID: " + astronaut.id);
      return;
    }
  }
  crew.push(astronaut);
}

addCrewMember(squad, firstAstronaut);

const remainingCrew = [
  { id: 2, name: "Bart", role: "Pilot", isEVAEligible: false, priority: 8 },
  { id: 3, name: "Caroline", role: "Engineer", isEVAEligible: true, priority: 4 },
  { id: 4, name: "Diego", role: "Scientist", isEVAEligible: false, priority: 1 },
  { id: 5, name: "Elise", role: "Medic", isEVAEligible: true, priority: 7 },
  { id: 6, name: "Felix", role: "Navigator", isEVAEligible: true, priority: 6 },
  { id: 7, name: "Gertrude", role: "Communications", isEVAEligible: false, priority: 4 },
  { id: 8, name: "Hank", role: "Mechanic", isEVAEligible: true, priority: 2 },
  { id: 9, name: "Irene", role: "Specialist", isEVAEligible: true, priority: 5 },
  { id: 10, name: "Joan", role: "Technician", isEVAEligible: false, priority: 1 },
];

for (let i = 0; i < remainingCrew.length; i++) {
  addCrewMember(squad, remainingCrew[i]);
}

function swapCrewMembers(crew, fromIndex, toIndex) {
  if (
    fromIndex < 0 || 
    toIndex < 0 ||
    fromIndex >= crew.length ||
    toIndex >= crew.length
  ) {
    console.log("Invalid crew indices");
    return;
  }

  const updatedCrew = crew.slice();
  --fcc-editable-region--
  
  --fcc-editable-region--
}
```
