---
id: 5f0d4ab1b435f13ab6550052
title: Step 54
challengeType: 0
dashedName: step-54
---

# --description--

Next, you are going to add some new form `input` elements, so add another `fieldset` element directly below the current `fieldset` element.

# --hints--

Kipengele chako kipya cha `fieldset` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert(document.querySelectorAll('fieldset').length >= 2);
```

Unapaswa kuongeza tu tagi moja ya ufunguzi ya `fieldset`. Tafadhali ondoa ziada yoyote.

```js
assert(document.querySelectorAll('fieldset').length === 2);
```

Kipengele chako cha `fieldset` kinafaa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert(code.match(/<\/fieldset>/g).length >= 2);
```

Unapaswa kuongeza tu tagi moja ya kufunga ya `fieldset`. Tafadhali ondoa ziada yoyote.

```js
assert(code.match(/<\/fieldset>/g).length === 2);
```

Kipengele cha pili cha `fieldset` hakipaswi kuwekwa katika kipengee cha kwanza cha `fieldset`.

```js
const childrenOf1stFieldset = [
  ...document.querySelector('form > fieldset').children
];
const foundElems = childrenOf1stFieldset.filter((child) => {
  return child.nodeName === 'FIELDSET';
});
assert(foundElems.length === 0);
```

Vipengele vyote viwili vya `fieldset` vinapaswa kuwa juu ya sehemu ya maandishi. Viko nje ya utaratibu.

```js
const formChildren = document.querySelector('form')?.children;
assert.isTrue(
  formChildren?.[0]?.nodeName === 'FIELDSET' &&
  formChildren?.[1]?.nodeName === 'FIELDSET' &&
  formChildren?.[2]?.nodeName === 'INPUT' &&
  formChildren?.[2]?.getAttribute('type') === 'text'
);
```

Kipengele chako kipya cha `fieldset` lazima kiwe chini ya kipengele cha `fieldset` kilichopo. Unazo katika mpangilio mbaya.

```js
const fieldsetChildren = [...document.querySelectorAll('fieldset')].map(
  (elem) => elem.children
);
assert(fieldsetChildren[0].length > fieldsetChildren[1].length);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
      <section>
        <h2>Cat Form</h2>
        <form action="https://freecatphotoapp.com/submit-cat-photo">
--fcc-editable-region--
          <fieldset>
            <legend>Is your cat an indoor or outdoor cat?</legend>
            <label><input id="indoor" type="radio" name="indoor-outdoor" value="indoor"> Indoor</label>
            <label><input id="outdoor" type="radio" name="indoor-outdoor" value="outdoor"> Outdoor</label>
          </fieldset>

--fcc-editable-region--
          <input type="text" name="catphotourl" placeholder="cat photo URL" required>
          <button type="submit">Submit</button>
        </form>
      </section>
    </main>
  </body>
</html>
```

