---
id: 587d7dba367417b2b2512ba8
title: Controlla tutto o niente
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

A volte i pattern che vuoi cercare possono contenere parti che potrebbero esserci oppure no. Tuttavia, potrebbe essere importante verificarle comunque.

Puoi indicare la possibile presenza di un elemento con un punto interrogativo, `?`. Questo verifica la presenza di zero o uno dell'elemento precedente. Puoi pensare a questo simbolo come a un modo per dire che l'elemento precedente è opzionale.

Per esempio, ci sono piccole differenze tra l'inglese americano e quello britannico e puoi usare il punto interrogativo per far corrispondere entrambe le grafie.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

Entrambi gli usi del metodo `test` restituirebbero `true`.

# --instructions--

Modifica l'espressione regolare `favRegex` per far corrispondere sia la versione americana (`favorite`) sia quella britannica (`favourite`) della parola.

# --hints--

La tua espressione regolare dovrebbe usare il simbolo opzionale, `?`.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

La tua espressione regolare dovrebbe corrispondere alla stringa `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

La tua espressione regolare dovrebbe corrispondere alla stringa `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

La tua espressione regolare non dovrebbe corrispondere alla stringa `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
