---
id: 587d7db7367417b2b2512b9d
title: "Corrispondenza di pattern all'inizio di una stringa"
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

Le sfide precedenti hanno mostrato che le espressioni regolari possono essere usate per cercare più corrispondenze. Sono anche usate per cercare pattern in posizioni specifiche all'interno delle stringhe.

In una sfida precedente, hai usato il carattere accento circonflesso (`^`) all'interno di un insieme di caratteri per creare un insieme negato nella forma `[^thingsThatWillNotBeMatched]`. Fuori da un insieme di caratteri, l'accento circonflesso si usa per cercare pattern all'inizio delle stringhe.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

La prima chiamata `test` restituirebbe `true`, mentre la seconda restituirebbe `false`.

# --instructions--

Usa il carattere accento circonflesso in un'espressione regolare per trovare `Cal` solo all'inizio della stringa `rickyAndCal`.

# --hints--

La tua espressione regolare dovrebbe cercare la stringa `Cal` con la lettera maiuscola.

```js
assert(calRegex.source == '^Cal');
```

La tua espressione regolare non dovrebbe usare flag.

```js
assert(calRegex.flags == '');
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Cal` all'inizio della stringa.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

La tua espressione regolare non dovrebbe corrispondere alla stringa `Cal` nel mezzo di una stringa.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
