---
id: 587d7db5367417b2b2512b97
title: "Corrispondenza di numeri e lettere dell'alfabeto"
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

Usare il trattino (`-`) per corrispondere a un intervallo di caratteri non è limitato alle lettere. Funziona anche per corrispondere a un intervallo di numeri.

Per esempio, `/[0-5]/` corrisponde a qualsiasi numero tra `0` e `5`, inclusi `0` e `5`.

Inoltre, è possibile combinare un intervallo di lettere e numeri in un unico insieme di caratteri.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

Crea una singola espressione regolare che corrisponda a un intervallo di lettere tra `h` e `s` e a un intervallo di numeri tra `2` e `6`. Ricorda di includere le flag appropriate nell'espressione regolare.

# --hints--

La tua espressione regolare `myRegex` dovrebbe corrispondere a 17 elementi.

```js
assert(result.length == 17);
```

La tua espressione regolare `myRegex` dovrebbe usare la flag globale.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

La tua espressione regolare `myRegex` dovrebbe usare la flag case insensitive.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
