---
id: 587d7db5367417b2b2512b96
title: "Corrispondenza delle lettere dell'alfabeto"
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

Hai visto come puoi usare i <dfn>set di caratteri</dfn> per specificare un gruppo di caratteri da far corrispondere, ma è molto lungo da scrivere quando devi far corrispondere un ampio intervallo di caratteri (per esempio, ogni lettera dell'alfabeto). Fortunatamente, c’è una funzionalità integrata che rende tutto breve e semplice.

All’interno di un set di caratteri, puoi definire un intervallo di caratteri da far corrispondere usando il carattere trattino: `-`.

Per esempio, per far corrispondere le lettere minuscole `a` fino a `e` useresti `[a-e]`.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

In ordine, le tre chiamate `match` restituirebbero i valori `["cat"]`, `["bat"]` e `null`.

# --instructions--

Fai corrispondere tutte le lettere nella stringa `quoteSample`.

**Nota**: assicurati di far corrispondere sia le lettere maiuscole sia quelle minuscole.

# --hints--

La tua espressione regolare `alphabetRegex` dovrebbe far corrispondere 35 elementi.

```js
assert(result.length == 35);
```

La tua espressione regolare `alphabetRegex` dovrebbe usare la flag globale.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

La tua espressione regolare `alphabetRegex` dovrebbe usare la flag case insensitive.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
