---
id: 587d7db7367417b2b2512b9e
title: Corrispondenza di pattern alla fine di una stringa
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

Nell'ultima sfida, hai imparato a usare il carattere accento circonflesso per cercare pattern all'inizio delle stringhe. Esiste anche un modo per cercare pattern alla fine delle stringhe.

Puoi cercare la fine delle stringhe usando il carattere del dollaro `$` alla fine dell'espressione regolare.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

La prima chiamata `test` restituirebbe `true`, mentre la seconda restituirebbe `false`.

# --instructions--

Usa il carattere àncora (`$`) per far corrispondere la stringa `caboose` alla fine della stringa `caboose`.

# --hints--

Dovresti cercare `caboose` con il carattere àncora del dollaro `$` nella tua espressione regolare.

```js
assert(lastRegex.source == 'caboose$');
```

La tua espressione regolare non dovrebbe usare alcun flag.

```js
assert(lastRegex.flags == '');
```

Dovresti far corrispondere `caboose` alla fine della stringa `The last car on a train is the caboose`

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
