---
id: 587d7db9367417b2b2512ba5
title: Specifica il numero minimo e massimo di corrispondenze
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

Ricorda che usi il segno più `+` per cercare uno o più caratteri e l’asterisco `*` per cercare zero o più caratteri. Sono comodi ma a volte vuoi trovare un certo intervallo di pattern.

Puoi specificare il numero minimo e massimo di pattern con i <dfn>quantificatori</dfn>. I quantificatori si usano con le parentesi graffe (`{` e `}`). Metti due numeri tra le parentesi graffe — per il numero minimo e massimo di pattern.

Per esempio, per trovare solo la lettera `a` che appare tra `3` e `5` volte nella stringa `ah`, la tua espressione regolare sarebbe `/a{3,5}h/`.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

La prima chiamata `test` restituirebbe `true`, mentre la seconda restituirebbe `false`.

# --instructions--

Modifica l’espressione regolare `ohRegex` per farla corrispondere all’intera frase `Oh no` solo quando contiene da `3` a `6` lettere `h`.

# --hints--

La tua espressione regolare dovrebbe usare le parentesi graffe.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

La tua espressione regolare non dovrebbe corrispondere alla stringa `Ohh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Ohhh no`

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Ohhhh no`

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Ohhhhh no`

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Ohhhhhh no`

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

La tua espressione regolare non dovrebbe corrispondere alla stringa `Ohhhhhhh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
