---
id: 587d7db9367417b2b2512ba6
title: Specificare solo il numero minimo di occorrenze
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

Puoi specificare il numero minimo e massimo di pattern con i quantificatori usando le parentesi graffe. A volte vuoi specificare solo il numero minimo di pattern senza un limite massimo.

Per specificare solo il numero minimo di pattern, mantieni il primo numero seguito da una virgola.

Per esempio, per far corrispondere solo la stringa `hah` con la lettera `a` che appare almeno `3` volte, la tua espressione regolare sarà `/ha{3,}h/`.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

In ordine, le tre chiamate `test` restituiranno `true`, `false` e `true`.

# --instructions--

Modifica l'espressione regolare `haRegex` per far corrispondere la parola `Hazzah` solo quando ha quattro o più lettere `z`.

# --hints--

La tua espressione regolare dovrebbe usare le parentesi graffe.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

La tua espressione regolare non dovrebbe far corrispondere la stringa `Hazzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

La tua espressione regolare non dovrebbe far corrispondere la stringa `Hazzzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

La tua espressione regolare dovrebbe far corrispondere la stringa `Hazzzzah`

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

La tua espressione regolare dovrebbe far corrispondere la stringa `Hazzzzzah`

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

La tua espressione regolare dovrebbe far corrispondere la stringa `Hazzzzzzah`

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

La tua espressione regolare dovrebbe far corrispondere la stringa `Hazzah` con 30 lettere `z` al suo interno.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
