---
id: 587d7db6367417b2b2512b9a
title: Corrispondenza di caratteri che compaiono zero o più volte
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

L'ultima sfida ha usato il segno più `+` per cercare caratteri che compaiono una o più volte. C'è anche un'opzione che corrisponde a caratteri che compaiono zero o più volte.

Il carattere per farlo è l'asterisco o stella: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

In ordine, le tre chiamate `match` restituirebbero i valori `["goooooooo"]`, `["g"]` e `null`.

# --instructions--

Crea una espressione regolare `chewieRegex` che usi il carattere `*` per corrispondere a un carattere maiuscolo `A` immediatamente seguito da zero o più caratteri minuscoli `a` in `chewieQuote`. La tua espressione regolare non ha bisogno di flag o classi di caratteri e non dovrebbe corrispondere a nessuna delle altre citazioni.

# --hints--

La tua espressione regolare `chewieRegex` dovrebbe usare il carattere `*` per corrispondere a zero o più caratteri `a`.

```js
assert(/\*/.test(chewieRegex.source));
```

La tua espressione regolare dovrebbe corrispondere alla stringa `A` in `chewieQuote`.

```js
assert(result[0][0] === 'A');
```

La tua espressione regolare dovrebbe corrispondere alla stringa `Aaaaaaaaaaaaaaaa` in `chewieQuote`.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

La tua espressione regolare `chewieRegex` dovrebbe corrispondere a 16 caratteri in `chewieQuote`.

```js
assert(result[0].length === 16);
```

La tua espressione regolare non dovrebbe corrispondere a nessun carattere nella stringa `He made a fair move. Screaming about it can't help you.`

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

La tua espressione regolare non dovrebbe corrispondere a nessun carattere nella stringa `Let him have it. It's not wise to upset a Wookiee.`

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
