---
id: 587d7db6367417b2b2512b9b
title: trovare caratteri con matching lazy
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

Nelle espressioni regolari, un match <dfn>greedy</dfn> trova la parte più lunga possibile di una stringa che corrisponde al pattern regex e la restituisce come match. L'alternativa si chiama match <dfn>lazy</dfn>, che trova la parte più piccola possibile della stringa che soddisfa il pattern regex.

Puoi applicare la regex `/t[a-z]*i/` alla stringa `"titanic"`. Questa regex è fondamentalmente un pattern che inizia con `t`, termina con `i` e ha alcune lettere nel mezzo.

Le espressioni regolari sono per impostazione predefinita greedy, quindi il match restituirebbe `["titani"]`. Trova la sottostringa più grande possibile che corrisponde al pattern.

Tuttavia, puoi usare il carattere `?` per cambiarlo in matching lazy. `"titanic"` abbinato alla regex modificata `/t[a-z]*?i/` restituisce `["ti"]`.

**Nota:** Evita di fare il parsing di HTML con espressioni regolari, ma usare le espressioni regolari per trovare pattern in una stringa HTML va benissimo.

# --instructions--

Correggi la regex `/<.*>/` per restituire il tag HTML `<h1>` e non il testo `"<h1>Winter is coming</h1>"`. Ricorda che il carattere jolly `.` in una espressione regolare corrisponde a qualsiasi carattere.

# --hints--

La variabile `result` dovrebbe essere un array con `<h1>` dentro

```js
assert(result[0] == '<h1>');
```

`myRegex` dovrebbe usare il matching lazy

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

`myRegex` non dovrebbe includere la stringa `h1`

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
