---
id: 587d7db4367417b2b2512b91
title: Ignora il maiuscolo/minuscolo durante la corrispondenza
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

Finora hai visto le espressioni regolari per fare corrispondenze letterali di stringhe. Ma a volte potresti voler considerare anche le differenze di maiuscole e minuscole.

Il maiuscolo (o a volte lettera maiuscola) è la differenza tra lettere maiuscole e lettere minuscole. Esempi di lettere maiuscole sono `A`, `B` e `C`. Esempi di lettere minuscole sono `a`, `b` e `c`.

Puoi far corrispondere entrambe le forme usando quella che si chiama flag. Ci sono altre flag, ma qui ti concentrerai sulla flag che ignora il maiuscolo/minuscolo: la flag `i`. Puoi usarla aggiungendola alla fine dell’espressione regolare. Un esempio di utilizzo di questa flag è `/ignorecase/i`. Questa regex può corrispondere alle stringhe `ignorecase`, `igNoreCase` e `IgnoreCase`.

# --instructions--

Scrivi una regex `fccRegex` per far corrispondere `freeCodeCamp`, indipendentemente dal maiuscolo o minuscolo. La tua regex non dovrebbe corrispondere ad abbreviazioni o varianti con spazi.

# --hints--

La tua regex dovrebbe corrispondere alla stringa `freeCodeCamp`

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

La tua regex dovrebbe corrispondere alla stringa `FreeCodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

La tua regex dovrebbe corrispondere alla stringa `FreecodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

La tua regex dovrebbe corrispondere alla stringa `FreeCodecamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

La tua regex non dovrebbe corrispondere alla stringa `Free Code Camp`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

La tua regex dovrebbe corrispondere alla stringa `FreeCOdeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

La tua regex non dovrebbe corrispondere alla stringa `FCC`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

La tua regex dovrebbe corrispondere alla stringa `FrEeCoDeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

La tua regex dovrebbe corrispondere alla stringa `FrEeCodECamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

La tua regex dovrebbe corrispondere alla stringa `FReeCodeCAmp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
