---
id: 587d7db8367417b2b2512ba3
title: Corrispondenza degli spazi bianchi
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

Finora le sfide hanno riguardato la corrispondenza di lettere dell’alfabeto e numeri. Puoi anche trovare gli spazi bianchi o gli spazi tra le lettere.

Puoi cercare gli spazi bianchi usando `\s`, che è un `s` minuscolo. Questo pattern non corrisponde solo agli spazi bianchi, ma anche ai caratteri di ritorno a capo, tabulazione, avanzamento modulo e nuova riga. Puoi pensarlo come simile alla classe di caratteri `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

Questa chiamata `match` restituirebbe `[" ", " "]`.
# --instructions--

Modifica l’espressione regolare `countWhiteSpace` per cercare più caratteri di spazio bianco in una stringa.

# --hints--

La tua espressione regolare dovrebbe usare il flag globale.

```js
assert(countWhiteSpace.global);
```

La tua espressione regolare dovrebbe usare il carattere di abbreviazione `\s` per corrispondere a tutti i caratteri di spazio bianco.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

La tua espressione regolare dovrebbe trovare otto spazi nella stringa `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

La tua espressione regolare dovrebbe trovare tre spazi nella stringa `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

La tua espressione regolare non dovrebbe trovare spazi nella stringa `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
