---
id: 587d7db9367417b2b2512ba4
title: trovare caratteri non di spazio
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

Hai imparato a cercare gli spazi usando `\s`, con un `s` minuscolo. Puoi anche cercare tutto tranne gli spazi.

Cerca i caratteri non di spazio usando `\S`, che è un `s` maiuscolo. Questo pattern non corrisponderà a spazi, ritorni a capo, tabulazioni, avanzamenti modulo e caratteri di nuova riga. Puoi pensarlo come simile alla classe di caratteri `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

Il valore restituito dal metodo `.length` sarebbe `32`.

# --instructions--

Modifica l'espressione regolare `countNonWhiteSpace` per cercare più caratteri non di spazio in una stringa.

# --hints--

La tua espressione regolare dovrebbe usare la flag globale.

```js
assert(countNonWhiteSpace.global);
```

La tua espressione regolare dovrebbe usare il carattere abbreviato `\S` per corrispondere a tutti i caratteri non di spazio.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

La tua espressione regolare dovrebbe trovare 35 caratteri non spazi nella stringa `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

La tua espressione regolare dovrebbe trovare 23 caratteri non spazi nella stringa `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

La tua espressione regolare dovrebbe trovare 21 caratteri non spazi nella stringa `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
