---
id: 587d7db6367417b2b2512b99
title: Corrispondenza di caratteri che si ripetono una o più volte
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

A volte, devi trovare una corrispondenza per un carattere (o gruppo di caratteri) che appare una o più volte di seguito. Questo significa che compare almeno una volta e può essere ripetuto.

Puoi usare il carattere `+` per verificare se è così. Ricorda, il carattere o pattern deve essere presente consecutivamente. Cioè, il carattere deve ripetersi uno dopo l'altro.

Per esempio, `/a+/g` troverebbe una corrispondenza in `abc` e restituirebbe `["a"]`. A causa del `+`, troverebbe anche una singola corrispondenza in `aabc` e restituirebbe `["aa"]`.

Se invece controllasse la stringa `abab`, troverebbe due corrispondenze e restituirebbe `["a", "a"]` perché i caratteri `a` non sono consecutivi — c'è un `b` tra di loro. Infine, dato che non c'è nessun `a` nella stringa `bcd`, non troverebbe alcuna corrispondenza.

# --instructions--

Vuoi trovare corrispondenze quando la lettera `s` compare una o più volte in `Mississippi`. Scrivi un'espressione regolare che usi il segno `+`.

# --hints--

La tua espressione regolare `myRegex` dovrebbe usare il segno `+` per corrispondere a uno o più caratteri `s`.

```js
assert(/\+/.test(myRegex.source));
```

La tua espressione regolare `myRegex` dovrebbe corrispondere a 2 elementi.

```js
assert(result.length == 2);
```

La variabile `result` dovrebbe essere un array con due corrispondenze di `ss`

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
