---
id: 587d7db4367417b2b2512b93
title: trovare più di una corrispondenza
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

Finora, sei stato in grado di estrarre o cercare un pattern solo una volta.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

Qui `match` restituirebbe `["Repeat"]`.

Per cercare o estrarre un pattern più di una volta, puoi usare la flag di ricerca globale: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

E qui `match` restituisce il valore `["Repeat", "Repeat", "Repeat"]`

# --instructions--

Usando l'espressione regolare `starRegex`, trova ed estrai entrambe le parole `Twinkle` dalla stringa `twinkleStar`.

**Nota**  
Puoi avere flag multipli nella tua espressione regolare come `/search/gi`

# --hints--

La tua espressione regolare `starRegex` dovrebbe usare la flag globale `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

La tua espressione regolare `starRegex` dovrebbe usare la flag case insensitive `i`

```js
assert(starRegex.flags.match(/i/).length == 1);
```

La tua corrispondenza dovrebbe trovare entrambe le occorrenze della parola `Twinkle`

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

La tua corrispondenza `result` dovrebbe contenere due elementi.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
