---
id: 587d7db5367417b2b2512b95
title: Corrispondenza di un singolo carattere con più possibilità
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

Hai imparato come trovare corrispondenze letterali (`/literal/`) e il carattere jolly (`/./`). Questi sono gli estremi delle espressioni regolari, dove uno trova corrispondenze esatte e l’altro corrisponde a tutto. Ci sono opzioni che rappresentano un equilibrio tra i due estremi.

Puoi cercare un pattern letterale con una certa flessibilità usando le <dfn>classi di caratteri</dfn>. Le classi di caratteri ti permettono di definire un gruppo di caratteri da trovare inserendoli tra parentesi quadre (`[` e `]`).

Per esempio, vuoi trovare `bag`, `big` e `bug` ma non `bog`. Puoi creare l’espressione regolare `/b[aiu]g/` per farlo. `[aiu]` è la classe di caratteri che corrisponderà solo ai caratteri `a`, `i` o `u`.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

In ordine, le quattro chiamate `match` restituirebbero i valori `["big"]`, `["bag"]`, `["bug"]` e `null`.

# --instructions--

Usa una classe di caratteri con le vocali (`a`, `e`, `i`, `o`, `u`) nella tua espressione regolare `vowelRegex` per trovare tutte le vocali nella stringa `quoteSample`.

**Nota:** Assicurati di trovare sia le vocali maiuscole sia quelle minuscole.

# --hints--

Dovresti trovare tutte e 25 le vocali.

```js
assert(result.length == 25);
```

La tua espressione regolare `vowelRegex` dovrebbe usare una classe di caratteri.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

La tua espressione regolare `vowelRegex` dovrebbe usare la flag globale.

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

La tua espressione regolare `vowelRegex` dovrebbe usare la flag case insensitive.

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

La tua espressione regolare non dovrebbe trovare nessuna consonante.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
