---
id: 587d7db5367417b2b2512b94
title: Corrispondenza di qualsiasi cosa con il punto jolly
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

A volte non conosci (o non hai bisogno di conoscere) esattamente i caratteri nei tuoi pattern. Pensare a tutte le parole che corrispondono, per esempio, a un errore di ortografia richiederebbe molto tempo. Per fortuna, puoi risparmiare tempo usando il carattere jolly: `.`

Il carattere jolly `.` corrisponderà a un qualsiasi carattere singolo. Il carattere jolly è anche chiamato `dot` e `period`. Puoi usare il carattere jolly proprio come qualsiasi altro carattere nell’espressione regolare. Per esempio, se vuoi far corrispondere `hug`, `huh`, `hut` e `hum`, puoi usare l’espressione regolare `/hu./` per far corrispondere tutte e quattro le parole.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Entrambe queste chiamate `test` restituirebbero `true`.

# --instructions--

Completa l’espressione regolare `unRegex` in modo che corrisponda alle stringhe `run`, `sun`, `fun`, `pun`, `nun` e `bun`. La tua espressione regolare dovrebbe usare il carattere jolly.

# --hints--

Dovresti usare il metodo `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Dovresti usare il carattere jolly nella tua espressione regolare `unRegex`

```js
assert(/\./.test(unRegex.source));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `run` nella stringa `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `sun` nella stringa `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `fun` nella stringa `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `pun` nella stringa `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `nun` nella stringa `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

La tua espressione regolare `unRegex` dovrebbe corrispondere a `bun` nella stringa `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

La tua espressione regolare `unRegex` non dovrebbe corrispondere alla stringa `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

La tua espressione regolare `unRegex` non dovrebbe corrispondere alla stringa `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
