---
id: 587d7db6367417b2b2512b98
title: Corrispondenza di singoli caratteri non specificati
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

Finora hai creato un insieme di caratteri che vuoi far corrispondere, ma potresti anche creare un insieme di caratteri che non vuoi far corrispondere. Questi tipi di insiemi di caratteri si chiamano <dfn>insiemi di caratteri negati</dfn>.

Per creare un insieme di caratteri negato, devi mettere un carattere accento circonflesso (`^`) subito dopo la parentesi di apertura e prima dei caratteri che non vuoi far corrispondere.

Per esempio, `/[^aeiou]/gi` corrisponde a tutti i caratteri che non sono vocali. Nota che caratteri come `.`, `!`, `[`, `@`, `/` e gli spazi bianchi sono corrisposti — l’insieme di caratteri vocalici negati esclude solo i caratteri vocalici.

# --instructions--

Crea una singola espressione regolare che corrisponda a tutti i caratteri che non sono un numero o una vocale. Ricorda di includere i flag appropriati nell’espressione regolare.

# --hints--

La tua espressione regolare `myRegex` dovrebbe corrispondere a 9 elementi.

```js
assert(result.length == 9);
```

La tua espressione regolare `myRegex` dovrebbe usare il flag globale.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

La tua espressione regolare `myRegex` dovrebbe usare il flag case insensitive.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
