---
id: 587d7db3367417b2b2512b8e
title: Usare il metodo test
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

Le espressioni regolari sono usate nei linguaggi di programmazione per trovare parti di stringhe. Crei dei pattern per aiutarti a fare questo confronto.

Se vuoi trovare la parola `the` nella stringa `The dog chased the cat`, potresti usare la seguente espressione regolare: `/the/`. Nota che le virgolette non sono necessarie all’interno dell’espressione regolare.

JavaScript ha diversi modi per usare le espressioni regolari. Un modo per testare una regex è usare il metodo `.test()`. Il metodo `.test()` prende la regex, la applica a una stringa (che va messa dentro le parentesi tonde) e restituisce `true` o `false` se il tuo pattern trova qualcosa oppure no.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

Il metodo `test` qui restituisce `true`.

# --instructions--

Applica la regex `myRegex` sulla stringa `myString` usando il metodo `.test()`.

# --hints--

Dovresti usare `.test()` per testare la regex.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Il risultato dovrebbe restituire `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
