---
id: 5c3dda8b4d8df89bea71600f
title: Controlla il raggruppamento misto di caratteri
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

A volte vogliamo controllare gruppi di caratteri usando un'espressione regolare e per farlo usiamo le parentesi tonde `()`.

Se vuoi trovare `Penguin` oppure `Pumpkin` in una stringa, puoi usare la seguente espressione regolare: `/P(engu|umpk)in/g`

Poi verifica se i gruppi di stringhe desiderati sono nella stringa di test usando il metodo `test()`.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

Il metodo `test` qui restituirebbe `true`.

# --instructions--

Correggi l'espressione regolare in modo che controlli i nomi `Franklin Roosevelt` oppure `Eleanor Roosevelt` in modo case sensitive e che permetta nomi di mezzo.

Poi correggi il codice in modo che l'espressione regolare che hai creato venga testata su `myString` e venga restituito `true` oppure `false` a seconda che l'espressione regolare corrisponda o meno.

# --hints--

La tua espressione regolare `myRegex` dovrebbe restituire `true` per la stringa `Franklin D. Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

La tua espressione regolare `myRegex` dovrebbe restituire `true` per la stringa `Eleanor Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

La tua espressione regolare `myRegex` dovrebbe restituire `false` per la stringa `Franklin Rosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

La tua espressione regolare `myRegex` dovrebbe restituire `false` per la stringa `Frank Roosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

La tua espressione regolare `myRegex` dovrebbe restituire `false` per la stringa `FranklinRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

La tua espressione regolare `myRegex` dovrebbe restituire `false` per la stringa `EleanorRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

Dovresti usare `.test()` per testare l'espressione regolare.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Il tuo risultato dovrebbe restituire `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
