---
id: 669e81368e52b3a5c35a2dc5
title: ビデオコンピレーションページを作成する
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `main` 要素の唯一の子として `body` 要素があること。
1. ページのトピックを示す `h1` 要素があること。
1. `h1` 要素の下に、ページのトピックを紹介する段落があること。
1. 最初の段落の下に3つの `section` 要素があること。
1. 各セクションには、順に `h2` 要素、段落、`iframe` 要素が含まれていること。
1. 3つの `iframe` 要素には、有効なビデオを指定する `src` 属性が設定されていること。
1. 各 `iframe` 要素には、埋め込まれたコンテンツを説明する `title` 属性と、適切なサイズに設定するための `height` 属性および `width` 属性が設定されていること。

# --hints--

`main` 要素の中に `body` 要素があるはずです。

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

`main` 要素の唯一の子として `body` 要素があるはずです。

```js
assert.equal(document.querySelector('body').children.length, 1);
```

`h1` 要素の中にページのトピックを示す `main` 要素があるはずです。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

`h1` 要素の下にページのトピックを紹介する段落があるはずです。

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

最初の `section` 要素の下に3つの `p` 要素があるはずです。

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

各 `section` 要素は、そのセクションのタイトルとして機能する `h2` 要素で始まるはずです。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

各 `section` 要素は、2番目の子としてビデオコンテンツを紹介する `p` 要素を含むはずです。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

各 `section` 要素は、3番目の子として `iframe` 要素を含むはずです。

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

最初の `iframe` 要素には、有効なビデオを指定する `src` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

最初の `iframe` 要素には、埋め込まれたコンテンツを説明する `title` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

最初の `iframe` 要素には `height` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

最初の `iframe` 要素には `width` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

2番目の `iframe` 要素には、有効なビデオを指定する `src` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

2番目の `iframe` 要素には、埋め込まれたコンテンツを説明する `title` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

2番目の `iframe` 要素には `height` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

2番目の `iframe` 要素には `width` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

3番目の `iframe` 要素には、有効なビデオを指定する `src` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

3番目の `iframe` 要素には、埋め込まれたコンテンツを説明する `title` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

3番目の `iframe` 要素には `height` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

3番目の `iframe` 要素には `width` 属性があるはずです。

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
