---
id: a77dbc43c33f39daa4429b4f
title: 构建一个布尔查看函数
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

在本实验中，你将构建一个函数，用于查看一个值是否被归类为布尔原语。

基本类型中的布尔值为 `true` 或者 `false`。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `booWho` 的函数，它接收一个参数。
1. 如果接收到的参数是布尔原语，函数应返回 `true`。
1. 如果参数是任何其他值，函数应返回 `false`。

# --hints--

你应该有一个 `booWho` 函数。

```js
assert.isFunction(booWho);
```

`booWho(true)` 应返回 `true`。

```js
assert.isTrue(booWho(true));
```

`booWho(false)` 应该返回 `true`。

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` 应该返回 `false`。

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` 应该返回 `false`。

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` 应该返回 `false`。

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` 应该返回 `false`。

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` 应该返回 `false`。

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` 应该返回 `false`。

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` 应该返回 `false`。

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` 应该返回 `false`。

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
