---
id: 69f0a37c3bf194a7c9c62492
title: デバイス貸出台帳を作成する
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

あなたはソフトウェア開発者として、会社のサービスデスクが管理するITハードウェアデバイスを管理する台帳アプリケーションを作成する任務を負っています。台帳は資産タグをキーとし、その値がオブジェクトであるオブジェクトです。

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. 台帳は、各デバイスが資産タグをキーに持ち、`type`、`status`、`borrower`（`name`と`email`を含む）、および`dueDate`プロパティを持つオブジェクトとしてモデル化してください。

1. `checkoutDevice(ledger, assetTag, borrower)`という関数を実装してください。この関数は台帳オブジェクトを受け取り、元の台帳を変更せずにクローンを作成します。クローン上で`borrower`オブジェクトの`name`と`email`を設定し、`status`を`"CheckedOut"`に更新し、`{ ledger: updatedLedger, message: confirmationString }`の形式のオブジェクトを返します。`message`には資産タグと`borrower`オブジェクトの`name`が含まれます。

   - `assetTag`が台帳に存在しない場合、関数は台帳を変更せずに返し、資産タグを含み見つからなかったことを説明する`message`を返してください。
   - デバイスがすでに`"CheckedOut"`の場合、関数は台帳を変更せずに返し、現在の借用者をそのままにし、資産タグを含みデバイスがすでに貸出中であることを説明する`message`を返してください。

1. `checkinDevice(ledger, assetTag)`という関数を実装してください。この関数は台帳オブジェクトを受け取り、元の台帳を変更せずにクローンを作成します。クローン上で`borrower`オブジェクトの`name`と`email`をクリアし、`dueDate`を`""`にリセットし、`status`を`"CheckedIn"`に更新し、`{ ledger: updatedLedger, message: confirmationString }`の形式のオブジェクトを返します。`message`には資産タグが含まれます。

   - `assetTag`が台帳に存在しない場合、関数は台帳を変更せずに返し、資産タグを含み見つからなかったことを説明する`message`を返してください。

1. `listOverdueDevices(ledger, today)`という関数を実装してください。この関数は台帳オブジェクトを受け取り、期限切れのデバイスオブジェクトの配列を返します。配列は`dueDate`の昇順でソートされます。デバイスが期限切れとは、`status`が`"CheckedOut"`であり、`dueDate`が`today`よりも厳密に前の日付である場合です。`today`と`dueDate`は月/日/年の順の文字列で、月と日はゼロ埋めされている場合とされていない場合があります（例：`"9/5/2025"`と`"09/05/2025"`の両方が有効です）。

1. コード内でJavaScriptの`Date`オブジェクトを使用しないでください。

1. `serializeLedger(ledger)`という関数を実装してください。この関数は台帳オブジェクトをJSON文字列に変換します。また、`loadLedger(json)`という関数を実装してください。この関数はJSON文字列をJavaScriptオブジェクトに変換します。

# --hints--

`checkoutDevice`という名前の関数を作成してください。

```js
assert.isFunction(checkoutDevice);
```

`checkinDevice`という名前の関数を作成してください。

```js
assert.isFunction(checkinDevice);
```

`listOverdueDevices`という名前の関数を作成してください。

```js
assert.isFunction(listOverdueDevices);
```

`serializeLedger`という名前の関数を作成してください。

```js
assert.isFunction(serializeLedger);
```

`loadLedger`という名前の関数を作成してください。

```js
assert.isFunction(loadLedger);
```

`checkoutDevice`関数は元の台帳を変更してはいけません。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

`checkoutDevice`関数は借用者の`name`と`email`を設定し、`status`を`"CheckedOut"`に設定してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

`checkoutDevice`関数は資産タグと借用者の`name`を含む確認メッセージを返してください。

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

`checkoutDevice`関数は関連しないデバイスや変更されていないフィールドを返される台帳で保持してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

`checkoutDevice`関数は資産タグが見つからなかった場合、台帳を変更せずに返し、資産タグを含むメッセージを返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

`checkoutDevice`関数はデバイスがすでに貸出中の場合、台帳を変更せずに返し、現在の借用者をそのままにし、資産タグを含むメッセージを返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

`checkinDevice`関数は元の台帳を変更してはいけません。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

`checkinDevice`関数は借用者の`name`と`email`をクリアし、`dueDate`を`""`にリセットし、`status`を`"CheckedIn"`に設定してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

`checkinDevice`関数は資産タグを含む確認メッセージを返してください。

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

`checkinDevice`関数は関連しないデバイスや変更されていないフィールドを返される台帳で保持してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

`checkinDevice`関数は資産タグが見つからなかった場合、台帳を変更せずに返し、資産タグを含むメッセージを返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

`listOverdueDevices`関数は配列を返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

`listOverdueDevices`関数は`status`が`"CheckedOut"`の期限切れデバイスのみを返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

`listOverdueDevices`関数は期限切れデバイスを`dueDate`の昇順で返してください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

`listOverdueDevices`関数は月や日がゼロ埋めされていなくても正しくソートしてください。

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

コード内でJavaScriptの`Date`オブジェクトを使用しないでください。

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

`serializeLedger`関数は渡された台帳を表すJSON文字列を返してください。

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

`loadLedger`関数はJSON文字列が表すJavaScriptオブジェクトを返してください。

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
