---
id: 5f331e55dfab7a896e53c3a1
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

Inside the `head` element, nest a `meta` element with an attribute named `charset` set to the value `utf-8`. This tells the browser how to encode characters for the page.

# --hints--

You should have a `meta` tag.

```js
assert.match(code, /<meta.*?\/?>/is);
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 标签应该有一个 `charset` 属性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 属性的值应该为 `utf-8`。

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
