---
id: 5f33294a6af5e9188dbdb8f3
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

It's time to add some menu content. Add a `main` element within the existing `body` element. It will eventually contain pricing information about coffee and desserts offered by the cafe.

# --hints--

你的代码应该有一个 `<main>` 开始标签。

```js
assert.match(code, /<main>/i);
```

你的代码应该有一个 `</main>` 结束标签。

```js
assert.match(code, /<\/main>/i);
```

你不应该改变你的 `body` 元素。 确认你没有不小心删除结束标签。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

你的 `main` 元素应该在你的 `body` 元素之后。

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
