---
id: 5f3313e74582ad9d063e3a38
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

Add a `head` element within the `html` element, so you can add a `title` element. The `title` element's text should be `Cafe Menu`.

# --hints--

你需要一个开始标签 `<head>`

```js
assert.match(code, /<head>/i);
```

需要一个结束标签 `</head>`

```js
assert.match(code, /<head>/i);
```

需要一个开始标签 `<title>`

```js
assert.match(code, /<title>/i);
```

需要一个结束标签 `</title>`

```js
assert.match(code, /<\/title>/i);
```

所有的 `title` 元素都应嵌套在 `head` 元素中。

```js
assert.match(code, /<head>\s*<title>.*<\/title>\s*<\/head>/si);
```

你的 `title` 元素内应该有文本 `Cafe Menu`。 您可能需要检查您的拼写。

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
