---
id: 5f332a88dc25a0fd25c7687a
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

The name of the cafe is `CAMPER CAFE`. So, add an `h1` element within your `main` element. Give it the name of the cafe in capitalized letters to make it stand out.

# --hints--

你应该有一个 `<h1>` 开始标签。

```js
assert(code.match(/<h1>/i));
```

你应该有一个 `</h1>` 结束标签。

```js
assert(code.match(/<\/h1>/i));
```

你不应该改变你现有的 `main` 元素。

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

你的 `h1` 元素应该嵌套在你的 `main` 元素中。

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

你的 `h1` 元素应该有文字 `CAMPER CAFE`。

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

