---
id: 699b6d91cde15459d3619c4a
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

就像 `profile` 对象一样，在控制台中查看 `userRoles` 对象会很有帮助。

添加一个 `console.log()` 来日志 `userRoles` 对象。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将 `userRoles` 日志到控制台。

```js
assert.equal(getLogs()[1], userRoles);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

--fcc-editable-region--

--fcc-editable-region--
```
