---
id: 699b62c99ec8c63470313365
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

现在是为你的 `profile` 对象添加另一个属性的时候了。

首先更新内联对象类型，添加一个 `bio` 属性，其类型设置为 `string`。

# --hints--

你应该更新内联对象类型，添加一个 `bio` 属性，类型设置为 `string`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
--fcc-editable-region--
  username: string;
  age: number;
  isLoggedIn: boolean;
  
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
