---
id: 699b65caf17a8e8d340af27d
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

现在是为 `profile` 对象添加最后一个属性的时候了。

在之前的课程中，你已经学习了如何使用数组类型。你了解了不同的数组类型，比如 `string[]`，或者 `number[]`：

```ts
const scores: number[] = [95, 87, 100];
```

在这个例子中，`scores` 数组只接受数字。如果你尝试添加其他类型，比如 `string`，或者 `null`，那么 TypeScript 会抛出误差。

更新你的 `profile` 对象，添加一个 `programmingLanguages` 属性，类型为 `string[]`。在对象字面中，将该属性设置为 `["JavaScript", "Python", "C++"]`。

# --hints--

你的 `programmingLanguages` 对象中应该有一个 `profile` 属性。

```js
assert.property(profile, "programmingLanguages");
```

你的 `programmingLanguages` 属性的值应该是 `["JavaScript", "Python", "C++"]`。

```js
assert.deepEqual(profile.programmingLanguages, ["JavaScript", "Python", "C++"]);
```

你的 `programmingLanguages` 属性应该被定义为 `string[]` 类型。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "programmingLanguages", type: "string[]" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};
--fcc-editable-region--

console.log(profile);
```
