---
id: 699b64096301fd75e3efeb1b
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

如果你再次打开控制台，你会看到不同的错误信息：

```md
Property 'bio' is missing in type '{ username: string; age: number; isLoggedIn: false; }' but required in type '{ username: string; age: number; isLoggedIn: boolean; bio: string; }'.
```

正如你在之前的课程中回顾的，你可以通过在属性名旁边使用 `?` 符号来将属性标记为可选的，像这样：

```ts
const musician: {
  name: string;
  instrument: string;
  isActive: boolean;
  numberOfAlbums?: number;
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

在这个例子中，`numberOfAlbums` 被标记为可选，因此 TypeScript 在这种情况下不会抛出误差。

将你的 `bio` 属性更新为标记为可选。

# --hints--

你的 `bio` 属性应该被标记为可选。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "bio", type: "string", isOptional: true };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
--fcc-editable-region--
  bio: string
--fcc-editable-region--
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
};

console.log(profile);
```
