---
id: 699b4e284291adcfcb90df47
title: 第一步
challengeType: 1
dashedName: step-1
---

# --description--

在本次工作坊中，你将通过创建一个类型安全的用户评测，练习使用原生类型以及对象和数组类型。

首先创建一个名为 `profile` 的变量，并赋值为一个对象。你的对象应包含以下三个属性及其对应的值：

```md
username: "codeLearner",
age: 25,
isLoggedIn: false,
```

# --hints--

你应该有一个名为 `profile` 的变量。

```js
assert.exists(profile);
```

你的 `profile` 变量应该是一个对象。

```js
assert.isObject(profile);
```

你的 `username` 属性应该在 `profile` 对象中。

```js
assert.property(profile, "username");
```

你的 `username` 属性的值应该是 `"codeLearner"`。

```js
assert.propertyVal(profile, "username", "codeLearner");
```

你的 `profile` 对象应该有一个 `age` 属性。

```js
assert.property(profile, "age");
```

你的 `age` 属性的值应该是 `25`。

```js
assert.propertyVal(profile, "age", 25);
```

你的 `profile` 对象应该有一个 `isLoggedIn` 属性。

```js
assert.property(profile, "isLoggedIn");
```

你的 `isLoggedIn` 属性的值应该是 `false`。

```js
assert.propertyVal(profile, "isLoggedIn", false);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```
