---
id: 699b68f94a0d154c2a7a5ce0
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

现在 `profile` 对象已经完成，是时候构建一个表现用户角色的新对象了。

创建一个名为 `userRoles` 的新变量，并赋值为一个对象。你的对象应包含以下属性和值：

```md
admin: "full-access",
editor: "limited-access",
viewer: "read-only"
```

# --hints--

你应该有一个名为 `userRoles` 的变量。

```js
assert.isDefined(userRoles);
```

你的 `userRoles` 变量应该是一个对象。

```js
assert.isObject(userRoles);
```

你的 `userRoles` 对象应该有一个 `admin` 属性。

```js
assert.property(userRoles, "admin");
```

你的 `admin` 属性应该有一个值 `"full-access"`。

```js
assert.propertyVal(userRoles, "admin", "full-access");
```

你的 `userRoles` 对象应该有一个 `editor` 属性。

```js
assert.property(userRoles, "editor");
```

你的 `editor` 属性应该有一个值 `"limited-access"`。

```js
assert.propertyVal(userRoles, "editor", "limited-access");
```

你的 `userRoles` 对象应该有一个 `viewer` 属性。

```js
assert.property(userRoles, "viewer");
```

你的 `viewer` 属性应该有一个值 `"read-only"`。

```js
assert.propertyVal(userRoles, "viewer", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--

--fcc-editable-region--
```
