---
id: 699b5370a43e299b4673bcaa
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

在构建 `profile` 对象时，将其属性显示在控制台会很有帮助。

为你的 `console.log` 添加一个将 `profile` 记录到控制台的日志。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该将 `profile` 记录到控制台。

```js
assert.equal(getLogs()[0], profile);
```

# --seed--

## --seed-contents--

```ts
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
}

--fcc-editable-region--

--fcc-editable-region--
```
