---
id: 699b6dd1a2961bd38491ef2e
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

`profile` 对象有严格的密钥和类型。如果你尝试添加未定义的属性，或者赋值不兼容的类型，TypeScript 会抛出误差。

如果 `userRoles` 对象不那么严格，密钥更灵活，那就好了。

正如你在之前的课程中回顾的，你可以使用 `Record` 类型来实现这个效果：

```ts
const studentGrades: Record<string, number> = {
  Math: 95,
  English: 88,
  Science: 92
};
```

在这个 `studentGrades` 示例中，TypeScript 不需要事先知道所有科目。这允许你以后添加更多科目，比如 `Art` 或 `Music`。TypeScript 只关心每个密钥对应的值是数字。

通过添加内联 `userRoles` 类型，更新你的 `Record` 对象。将其设置为 `<string, string>`，使每个密钥可以是任意字串，每个值必须是字串。

# --hints--

你的 `userRoles` 对象应该有一个内联 `Record` 类型，设置为 `Record<string, string>`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.annotation.matches("Record<string, string>"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles = {
--fcc-editable-region--
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"
};

console.log(userRoles);
```
