---
id: 699b61ca8c9cb9bc1968364e
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

如果你打开控制台，你应该会看到以下错误信息：

```md
Object literal may only specify known properties, and 'mood' does not exist in type '{ username: string; age: number; isLoggedIn: boolean; }'.
```

在原生 JavaScript 中，你可以随意为 `profile` 对象添加任何属性和值。但在 TypeScript 中，编译器会检查对象是否只包含你定义的属性（本例中是 `username`、`age` 和 `isLoggedIn`）。由于 `mood` 没有包含在类型中，TypeScript 会将其视为无效属性并抛出错误。

请从你的代码中删除 `mood: null`。

# --hints--

你的代码中不应再有 `mood: null`。

```js
assert.isUndefined(profile.mood);
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false
--fcc-editable-region--
  mood: null
--fcc-editable-region--
};

console.log(profile);
```
