---
id: 699b5b7cb68b6d35afca61b7
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

虽然可以继续像这样为 `profile` 对象添加属性，但最好限制允许的属性类型。预先定义 `profile` 对象的结构会更好。

在之前的课程中，你已经学习了如何使用这样的对象类型：

```ts
const composer: { 
  name: string; 
  instrument: string; 
  isActive: boolean; 
} = {
  name: "Ludwig van Beethoven",
  instrument: "Piano",
  isActive: false
};
```

在这个例子中，`composer` 是一个带有内联对象类型的对象。该对象有三个必需的属性，分别是 `name`、`instrument` 和 `isActive`。前两个属性必须是 `string` 类型，`isActive` 必须是 `boolean` 类型。

更新你现有的 `profile` 对象，使其包含一个内联对象类型。`username` 应该被定义为 `string` 类型，`age` 应该被定义为 `number` 类型，`isLoggedIn` 应该被定义为 `boolean` 类型。

# --hints--

你的 `profile` 对象应该有一个内联对象类型，其中 `username` 属性被设置为 `string`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "username", type: "string" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

你的 `profile` 对象应该有一个内联对象类型，其中 `age` 属性被设置为 `number`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "age", type: "number" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

你的 `profile` 对象应该有一个内联对象类型，其中 `isLoggedIn` 属性被设置为 `boolean`。

```js
const explorer = await __helpers.Explorer(code);
const { profile } = explorer.variables;
const prop = { name: "isLoggedIn", type: "boolean" };
assert.isTrue(profile.annotation.hasTypeProps(prop));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
const profile = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  mood: null
}
--fcc-editable-region--

console.log(profile);
```
