---
id: 699b71d98c820ac7ab9eea2a
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

如果你再次打开控制台，你会看到一个新的错误信息：

```md
Type 'number' is not assignable to type 'string'.
```

这是因为你试图为 `guest` 属性添加一个数字值。

要消除 TypeScript 误差，请将 `guest` 属性值更改为 `"read-only"`。

完成最后的更改后，你的用户评测工作坊就完成了！

# --hints--

你的 `guest` 属性对于 `userRoles` 对象应该有 `"read-only"` 的值。

```js
assert.propertyVal(userRoles, "guest", "read-only");
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
--fcc-editable-region--
  guest: 3
--fcc-editable-region--
};

console.log(userRoles);
```

# --solutions--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only",
  moderator: "medium-access",
  guest: "read-only"
};

console.log(userRoles);
```
