---
id: 699b70b10ae344c802d106b0
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

现在是为 `userRoles` 对象添加更多属性的时候了。

添加一个 `moderator` 属性，值为 `"medium-access"`。

然后在它下面，添加一个 `guest` 属性，值为 `3`。

# --hints--

你的 `userRoles` 对象应该有一个 `moderator` 属性。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.moderator);
```

你的 `moderator` 属性应该有一个值为 `"medium-access"`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.moderator.value.matches("'medium-access'"));
```

你的 `userRoles` 对象应该有一个 `guest` 属性。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.exists(userRoles.objectProps.guest);
```

你的 `guest` 属性应该有一个值为数字 `3`。

```js
const explorer = await __helpers.Explorer(code);
const { userRoles } = explorer.variables;
assert.isTrue(userRoles.objectProps.guest.value.matches("3"));
```

# --seed--

## --seed-contents--

```ts
const profile: {
  username: string;
  age: number;
  isLoggedIn: boolean;
  bio?: string;
  programmingLanguages: string[];
} = {
  username: "codeLearner",
  age: 25,
  isLoggedIn: false,
  programmingLanguages: ["JavaScript", "Python", "C++"]
};

console.log(profile);

--fcc-editable-region--
const userRoles: Record<string, string> = {
  admin: "full-access",
  editor: "limited-access",
  viewer: "read-only"

};
--fcc-editable-region--

console.log(userRoles);
```
