---
id: a3566b1109230028080c9345
title: Спроєктуйте алгоритм суми всіх чисел
challengeType: 26
dashedName: lab-sum-all-numbers-algorithm
---

# --description--

У цій лабораторній роботі вам потрібно спроєктувати алгоритм суми всіх чисел.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `sumAll`, яка приймає масив із двох чисел.
1. `sumAll([n, m])` має повертати суму `n` і `m` плюс суму всіх чисел між ними. Найменше число не завжди буде першим. Наприклад, `sumAll([4,1])` має повертати `10`, бо сума всіх чисел між `1` і `4` (включно) дорівнює `10`.

# --hints--

У вас має бути функція з назвою `sumAll`.

```js
assert.isFunction(sumAll);
```

`sumAll([1, 4])` має повертати число.

```js
assert.isNumber(sumAll([1, 4]));
```

`sumAll([1, 4])` має повертати `10`.

```js
assert.equal(sumAll([1, 4]), 10);
```

`sumAll([4, 1])` має повертати `10`.

```js
assert.equal(sumAll([4, 1]), 10);
```

`sumAll([5, 10])` має повертати `45`.

```js
assert.equal(sumAll([5, 10]), 45);
```

`sumAll([10, 5])` має повертати `45`.

```js
assert.equal(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function sumAll(arr) {
  let sum = 0;
  
  arr.sort((a, b) => a - b);

  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }

  return sum;
}
```
