---
id: 594810f028c0303b75339ad2
title: ベクトルのクロス積
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

ベクトルは3つの数値の順序付き集合 (X, Y, Z) で表される3次元のものとして定義されます。

# --instructions--

2つのベクトル（配列）を入力として受け取り、それらのクロス積を計算する関数を書いてください。異なる長さのベクトルなど無効な入力の場合は、`null`を返す関数にしてください。

# --hints--

`crossProduct`は関数であるべきです。

```js
assert.equal(typeof crossProduct, 'function');
```

`crossProduct()`はnullを返すべきです。

```js
assert.equal(crossProduct(), null);
```

`crossProduct([1, 2, 3], [4, 5, 6])`は`[-3, 6, -3]`を返すべきです。

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
