---
id: 594810f028c0303b75339acf
title: アッカーマン関数
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

アッカーマン関数は再帰関数の古典的な例であり、特に原始再帰関数ではないことで知られています。値が非常に速く増加し、その呼び出しツリーのサイズも同様に大きくなります。

アッカーマン関数は通常、次のように定義されます：

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

引数は決して負の値にならず、常に終了します。

# --instructions--

$A(m, n)$ の値を返す関数を書いてください。関数の値が非常に速く増加するため任意精度が望ましいですが、必須ではありません。

# --hints--

`ack` は関数であるべきです。

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` は 1 を返すべきです。

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` は 3 を返すべきです。

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` は 13 を返すべきです。

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` は 61 を返すべきです。

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
