---
id: 594810f028c0303b75339ad4
title: ワードラップ
challengeType: 1
forumTopicId: 302344
dashedName: word-wrap
---

# --description--

現在でも、プロポーショナルフォントや複雑なレイアウトがある中で、指定した列でテキストを折り返す必要がある場合があります。基本的な課題は、テキストの段落をシンプルに折り返すことです。

# --instructions--

任意の文字数でこのテキストを折り返せる関数を書いてください。入力テキストにはすでに改行が含まれているため、関数はそれを適切に処理する必要があります。例として、80文字で折り返したテキストは以下のようになります。

<pre>
Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX
algorithm. If your language provides this, you get easy extra credit, but you
must reference documentation indicating that the algorithm is something better
than a simple minimum length algorithm.
</pre>

# --before-each--

```js
const text =
`Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX algorithm.
If your language provides this, you get easy extra credit,
but you ''must reference documentation'' indicating that the algorithm
is something better than a simple minimum length algorithm.`;

const wrapped80 = wrap(text, 80);
const wrapped42 = wrap(text, 42);

const firstRow80 =
    'Wrap text using a more sophisticated algorithm such as the Knuth and Plass TeX';

const firstRow42 = 'Wrap text using a more sophisticated';
```

# --hints--

`wrap` は関数であるべきです。

```js
assert.equal(typeof wrap, 'function');
```

`wrap` は文字列を返すべきです。

```js
assert.equal(typeof wrap('abc', 10), 'string');
```

`wrap(text,80)` は4行を返すべきです。

```js
assert(wrapped80.split('\n').length === 4);
```

あなたの `wrap` 関数は期待されるテキストを返すべきです。

```js
assert.equal(wrapped80.split('\n')[0], firstRow80);
```

`wrap(text,42)` は7行を返すべきです。

```js
assert(wrapped42.split('\n').length === 7);
```

あなたの `wrap` 関数は期待されるテキストを返すべきです。

```js
assert.equal(wrapped42.split('\n')[0], firstRow42);
```

# --seed--

## --seed-contents--

```js
function wrap(text, limit) {
  return text;
}
```

# --solutions--

```js
function wrap(text, limit) {
  const noNewlines = text.replace('\n', '');
  if (noNewlines.length > limit) {
    // find the last space within limit
    const edge = noNewlines.slice(0, limit).lastIndexOf(' ');
    if (edge > 0) {
      const line = noNewlines.slice(0, edge);
      const remainder = noNewlines.slice(edge + 1);
      return line + '\n' + wrap(remainder, limit);
    }
  }
  return text;
}
```
