---
id: 594810f028c0303b75339acc
title: ABC問題
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

ABCブロックのセット（例えば、子供の頃のアルファベットブロック）が与えられます。20個のブロックがあり、それぞれのブロックには2つの文字が書かれています。すべてのブロックの面を合わせると完全なアルファベットが保証されています。ブロックのサンプルセット：

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

文字列（単語）を受け取り、その単語が与えられたブロックのセットで綴れるかどうかを判定する関数を実装してください。

覚えておくべきルール：

<ul>
  <li>ブロックの文字を一度使ったら、そのブロックは再度使えません。</li>
  <li>関数は大文字・小文字を区別しません。</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

`canMakeWord` は関数であるべきです。

```js
assert(typeof canMakeWord === 'function');
```

`canMakeWord` はブール値を返すべきです。

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

`canMakeWord("bark")` は true を返すべきです。

```js
assert(canMakeWord(words[0]));
```

`canMakeWord("BooK")` は false を返すべきです。

```js
assert(!canMakeWord(words[1]));
```

`canMakeWord("TReAT")` は true を返すべきです。

```js
assert(canMakeWord(words[2]));
```

`canMakeWord("COMMON")` は false を返すべきです。

```js
assert(!canMakeWord(words[3]));
```

`canMakeWord("squAD")` は true を返すべきです。

```js
assert(canMakeWord(words[4]));
```

`canMakeWord("conFUSE")` は true を返すべきです。

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
