---
id: 594810f028c0303b75339ace
title: アキュムレーターファクトリー
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

Paul Grahamが提起した問題は、単一の（数値の）引数を取り、別の関数を返す関数を作ることです。この返される関数はアキュムレーターであり、さらに単一の数値引数を取り、これまでにそのアキュムレーターに渡されたすべての数値の合計（アキュムレーターが作成されたときに渡された初期値も含む）を返します。

# --instructions--

数値$n$を受け取り、それまでに渡されたすべての数の合計を返すアキュムレーター関数を生成する関数を作成してください。

**ルール:**

グローバル変数は使用しないでください。

**ヒント:**

クロージャは外側の状態を保存します。

# --hints--

`accumulator`は関数であるべきです。

```js
assert(typeof accumulator === 'function');
```

`accumulator(0)`は関数を返すべきです。

```js
assert(typeof accumulator(0) === 'function');
```

`accumulator(0)(2)`は数値を返すべきです。

```js
assert(typeof accumulator(0)(2) === 'number');
```

3、-4、1.5、5の値を渡すと5.5を返すべきです。

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
