---
id: 594810f028c0303b75339ad5
title: Yコンビネーター
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

厳密な<a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow">関数型プログラミング</a>やラムダ計算では、関数（ラムダ式）はステートを持たず、囲んでいる関数の引数のみを参照できます。これにより、関数が変数の状態に関連付けられ、その変数の状態が関数本体で使われるという通常の再帰関数の定義は除外されます。

Yコンビネーターはステートレスな関数で、別のステートレスな関数に適用すると、その関数の再帰バージョンを返します。Yコンビネーターは、不動点コンビネーターと呼ばれるこの種の関数の中で最も単純なものです。

# --instructions--

ステートレスなYコンビネーター関数を定義し、それを使って階乗を計算してください。`factorial(N)`関数はすでに用意されています。

# --hints--

Yは関数を返す必要があります。

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

factorial(1)は1を返すべきです。

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

factorial(2)は2を返すべきです。

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

factorial(3)は6を返すべきです。

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

factorial(4)は24を返すべきです。

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

factorial(10)は3628800を返すべきです。

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
