---
id: 594810f028c0303b75339ad1
title: ハッピーナンバー
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

ハッピーナンバーは次の手順で定義されます：

任意の正の整数から始め、その数字の各桁の二乗の和に置き換え、この処理を繰り返します。数が `1` になったら（そのまま変わらず）、または `1` を含まないサイクルで無限にループするまで続けます。この処理が `1` で終わる数はハッピーナンバーであり、 `1` で終わらない数はアンハッピーナンバーです。

# --instructions--

数がハッピーナンバーなら true を、そうでなければ false を返す関数を実装してください。

# --hints--

`happy` は関数であるべきです。

```js
assert(typeof happy === 'function');
```

`happy(1)` はブール値を返すべきです。

```js
assert(typeof happy(1) === 'boolean');
```

`happy(1)` は `true` を返すべきです。

```js
assert(happy(1));
```

`happy(2)` は `false` を返すべきです。

```js
assert(!happy(2));
```

`happy(7)` は `true` を返すべきです。

```js
assert(happy(7));
```

`happy(10)` は `true` を返すべきです。

```js
assert(happy(10));
```

`happy(13)` は `true` を返すべきです。

```js
assert(happy(13));
```

`happy(19)` は `true` を返すべきです。

```js
assert(happy(19));
```

`happy(23)` は `true` を返すべきです。

```js
assert(happy(23));
```

`happy(28)` は `true` を返すべきです。

```js
assert(happy(28));
```

`happy(31)` は `true` を返すべきです。

```js
assert(happy(31));
```

`happy(32)` は `true` を返すべきです。

```js
assert(happy(32));
```

`happy(33)` は `false` を返すべきです。

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
