---
id: 594810f028c0303b75339acb
title: 100のドア
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

100個のドアが一列に並んでいて、すべて最初は閉まっています。ドアの前を100回通ります。最初の通過では、すべてのドアを訪れて「トグル」します（ドアが閉まっていれば開け、開いていれば閉めます）。2回目の通過では、2番目ごとのドア（つまりドア#2、#4、#6、…）だけを訪れてトグルします。3回目は3番目ごとのドア（ドア#3、#6、#9、…）を訪れてトグルします。これを繰り返し、最後は100番目のドアだけを訪れます。

# --instructions--

最後の通過後のドアの状態を判定する関数を実装してください。結果は配列で返し、開いているドアの番号だけを配列に含めてください。

# --hints--

`getFinalOpenedDoors` は関数である必要があります。

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` は配列を返す必要があります。

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` は正しい結果を出す必要があります。

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
