---
id: 594810f028c0303b75339acd
title: '過剰数、不足数、完全数の分類'
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

これらは真の約数に基づく正の整数の3つの分類を定義します。

$P(n)$ を `n` の真の約数の合計とし、真の約数は `n` で `n` 自身を除くすべての正の整数です。

もし `P(n) < n` ならば `n` は `deficient` に分類されます。

もし `P(n) === n` ならば `n` は `perfect` に分類されます。

もし `P(n) > n` ならば `n` は `abundant` に分類されます。

**例**: `6` の真の約数は `1`、`2`、および `3` です。`1 + 2 + 3 = 6` なので `6` は完全数に分類されます。

# --instructions--

`1` から `num` まで（両端を含む）の整数のうち、3つの分類のそれぞれに属する数がいくつあるかを計算する関数を実装してください。結果は `[deficient, perfect, abundant]` の形式の配列で出力してください。

# --hints--

`getDPA` は関数であるべきです。

```js
assert(typeof getDPA === 'function');
```

`getDPA(5000)` は配列を返すべきです。

```js
assert(Array.isArray(getDPA(5000)));
```

`getDPA(5000)` の返す配列の長さは `3` であるべきです。

```js
assert(getDPA(5000).length === 3);
```

`getDPA(5000)` は `[3758, 3, 1239]` を返すべきです。

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

`getDPA(10000)` は `[7508, 4, 2488]` を返すべきです。

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

`getDPA(20000)` は `[15043, 4, 4953]` を返すべきです。

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
