---
id: 5949b579404977fbaefcd737
title: 友愛数ペア
challengeType: 1
forumTopicId: 302225
dashedName: amicable-pairs
---

# --description--

2つの整数 $N$ と $M$ は、$N \\neq M$ であり、$N$ の真の約数の和 ($\\mathrm{sum}(\\mathrm{propDivs}(N))$) が $M$ に等しく、かつ $\\mathrm{sum}(\\mathrm{propDivs}(M)) = N$ である場合、友愛数ペアと呼ばれます。

**例：**

**1184** と **1210** は友愛数ペアであり、それぞれの真の約数は以下の通りです：

<ul>
  <li>1、2、4、8、16、32、37、74、148、296、592</li>
  <li>1、2、5、10、11、22、55、110、121、242、605</li>
</ul>

最初の値 **1184** の約数の和は **1210** であり、2番目の値 **1210** の約数の和は **1184** です。

# --instructions--

20,000未満の友愛数ペアを計算してここに表示してください（8組あります）。

# --before-each--

```js
const answer300 = [[220, 284]];
const answer3000 = [
  [220, 284],
  [1184, 1210],
  [2620, 2924]
];
const answer20000 = [
  [220, 284],
  [1184, 1210],
  [2620, 2924],
  [5020, 5564],
  [6232, 6368],
  [10744, 10856],
  [12285, 14595],
  [17296, 18416]
];
```

# --hints--

`amicablePairsUpTo` は関数である必要があります。

```js
assert(typeof amicablePairsUpTo === 'function');
```

`amicablePairsUpTo(300)` は `[[220,284]]` を返す必要があります。

```js
assert.deepEqual(amicablePairsUpTo(300), answer300);
```

`amicablePairsUpTo(3000)` は `[[220,284],[1184,1210],[2620,2924]]` を返す必要があります。

```js
assert.deepEqual(amicablePairsUpTo(3000), answer3000);
```

`amicablePairsUpTo(20000)` は `[[220,284],[1184,1210],[2620,2924],[5020,5564],[6232,6368],[10744,10856],[12285,14595],[17296,18416]]` を返す必要があります。

```js
assert.deepEqual(amicablePairsUpTo(20000), answer20000);
```

# --seed--

## --seed-contents--

```js
function amicablePairsUpTo(maxNum) {

  return true;
}
```

# --solutions--

```js
// amicablePairsUpTo :: Int -> [(Int, Int)]
function amicablePairsUpTo(maxNum) {
  return range(1, maxNum)
    .map(x => properDivisors(x)
      .reduce((a, b) => a + b, 0))
    .reduce((a, m, i, lst) => {
      const n = i + 1;

      return (m > n) && lst[m - 1] === n ?
        a.concat([
          [n, m]
        ]) : a;
    }, []);
}

// properDivisors :: Int -> [Int]
function properDivisors(n) {
  if (n < 2) return [];

  const rRoot = Math.sqrt(n);
  const intRoot = Math.floor(rRoot);
  const blnPerfectSquare = rRoot === intRoot;
  const lows = range(1, intRoot)
  .filter(x => (n % x) === 0);

  return lows.concat(lows.slice(1)
    .map(x => n / x)
    .reverse()
    .slice(blnPerfectSquare | 0));
}

// Int -> Int -> Maybe Int -> [Int]
function range(m, n, step) {
  const d = (step || 1) * (n >= m ? 1 : -1);

  return Array.from({
    length: Math.floor((n - m) / d) + 1
  }, (_, i) => m + (i * d));
}
```
