---
id: 5949b579404977fbaefcd736
title: 神の90億の名前：整数編
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

この課題はアーサー・C・クラークの短編小説のバリエーションです。

（この課題を完了することの影響を理解しておいてください。）

「名前」とは何を意味するかを詳しく説明すると：

<ul>
  <li>整数1は1つの名前「1」を持ちます。</li>
  <li>整数2は2つの名前「1+1」と「2」を持ちます。</li>
  <li>整数3は3つの名前「1+1+1」、「2+1」、および「3」を持ちます。</li>
  <li>整数4は5つの名前「1+1+1+1」、「2+1+1」、「2+2」、「3+1」、「4」を持ちます。</li>
  <li>整数5は7つの名前「1+1+1+1+1」、「2+1+1+1」、「2+2+1」、「3+1+1」、「3+2」、「4+1」、「5」を持ちます。</li>
</ul>

これは次のように視覚化できます：

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

ここで、行 $n$ は整数 $n$ に対応し、行 $m$ の左から右への各列 $C$ は $C$ で始まる名前の数に対応します。

オプションとして、$n$ 行目の合計 $P(n)$ は整数分割関数であることに注意してください。

# --instructions--

$n$ 行目の合計を返す関数を実装してください。

# --hints--

`numberOfNames` は関数であるべきです。

```js
assert(typeof numberOfNames === 'function');
```

`numberOfNames(5)` は7に等しいべきです。

```js
assert.equal(numberOfNames(5), 7);
```

`numberOfNames(12)` は77に等しいべきです。

```js
assert.equal(numberOfNames(12), 77);
```

`numberOfNames(18)` は385に等しいべきです。

```js
assert.equal(numberOfNames(18), 385);
```

`numberOfNames(23)` は1255に等しいべきです。

```js
assert.equal(numberOfNames(23), 1255);
```

`numberOfNames(42)` は53174に等しいべきです。

```js
assert.equal(numberOfNames(42), 53174);
```

`numberOfNames(123)` は2552338241に等しいべきです。

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
