---
id: 594dc6c729e5700999302b45
title: 括弧のバランス
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

生成された括弧の文字列がバランスしているかどうかを判定してください。つまり、すべてが開き括弧と閉じ括弧のペア（その順番）で構成されており、誤ったネストがないかどうかを確認します。

**例:**
| 入力                       | 出力   |
| ------------------------- | ------ |
| <code>\[]</code>          | true   |
| <code>]\[</code>          | false  |
| <code>[][]</code>         | true   |
| <code>]\[]</code>         | false  |
| <code>\[]]\[\[]</code>    | false  |
| <code>\[\[\[\[]]]]</code> | true   |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

`isBalanced` は関数である必要があります。

```js
assert(typeof isBalanced === 'function');
```

`isBalanced("[]")` は true を返す必要があります。

```js
assert(isBalanced(testCases[0]));
```

`isBalanced("]][[[][][][]][")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[1]));
```

`isBalanced("[][[[[][][[[]]]]]]")` は true を返す必要があります。

```js
assert(isBalanced(testCases[2]));
```

`isBalanced("][")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[3]));
```

`isBalanced("[[[]]]][[]")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[4]));
```

`isBalanced("][[]")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[5]));
```

`isBalanced("][[][]][[[]]")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[6]));
```

`isBalanced("[[][]]][")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[7]));
```

`isBalanced("[[[]]][[]]]][][[")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[8]));
```

`isBalanced("[]][[]]][[[[][]]")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[9]));
```

`isBalanced("][]][[][")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[10]));
```

`isBalanced("[[]][[][]]")` は true を返す必要があります。

```js
assert(isBalanced(testCases[11]));
```

`isBalanced("[[]]")` は true を返す必要があります。

```js
assert(isBalanced(testCases[12]));
```

`isBalanced("]][]][[]][[[")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[13]));
```

`isBalanced("][]][][[")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[14]));
```

`isBalanced("][][")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[15]));
```

`isBalanced("[]]]")` は false を返す必要があります。

```js
assert(!isBalanced(testCases[16]));
```

`isBalanced("")` は true を返す必要があります。

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
