---
id: 594d8d0ab97724821379b1e6
title: 平均値と最頻値
challengeType: 1
forumTopicId: 302226
dashedName: averagesmode
---

# --description--

配列の中で最も頻繁に現れる値を見つける関数`mode`を書いてください。

コレクションが空の場合は無視してかまいません。最頻値が一意でない場合の処理に注意してください。

一般的なコレクションをサポートするのが適切でない、または不可能な場合は、可能であればベクター（配列）を使用してください。指定されていない値の型をサポートするのが適切でない、または不可能な場合は、整数を使用してください。

# --before-each--

```js
const arr1 = [1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17];
const arr2 = [1, 2, 4, 4, 1];
```

# --hints--

`mode`は関数であるべきです。

```js
assert(typeof mode === 'function');
```

`mode([1, 3, 6, 6, 6, 6, 7, 7, 12, 12, 17])`は`[6]`と等しくなるべきです。

```js
assert.deepEqual(mode(arr1), [6]);
```

`mode([1, 2, 4, 4, 1])`は`[1, 4]`と等しくなるべきです。

```js
assert.deepEqual(mode(arr2).sort(), [1, 4]);
```

# --seed--

## --seed-contents--

```js
function mode(arr) {

  return true;
}
```

# --solutions--

```js
function mode(arr) {
  const counter = {};
  let result = [];
  let max = 0;
  // for (const i in arr) {
  arr.forEach(el => {
    if (!(el in counter)) {
      counter[el] = 0;
    }
    counter[el]++;

    if (counter[el] === max) {
      result.push(el);
    }
    else if (counter[el] > max) {
      max = counter[el];
      result = [el];
    }
  });
  return result;
}
```
