---
id: 594ecc0d9a8cf816e3340187
title: タクシーキャブ数
challengeType: 1
forumTopicId: 302337
dashedName: taxicab-numbers
---

# --description--

タクシーキャブ数（ここで使われている定義）は、2つの正の立方数の和として複数の方法で表せる正の整数です。

最初のタクシーキャブ数は `1729` で、これは次のように表せます：

1<sup>3</sup> + 12<sup>3</sup> および

9<sup>3</sup> + 10<sup>3</sup>。

タクシーキャブ数は次のようにも呼ばれます：

<ul>
  <li>taxi numbers</li>
  <li>taxi-cab numbers</li>
  <li>taxi cab numbers</li>
  <li>Hardy-Ramanujan numbers</li>
</ul>

# --instructions--

最も小さい `n` 個のタクシーキャブ数を返す関数を書いてください。各タクシーキャブ数について、その数値と構成する立方数も表示してください。

# --before-each--

```js
const res4 = [1729, 4104, 13832, 20683];
const res25 = [
  1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656,
  110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763,
  373464, 402597
];

const res39From20To29 = [314496, 320264, 327763, 373464, 402597, 439101, 443889, 513000, 513856];
```

# --hints--

`taxicabNumbers` は関数である必要があります。

```js
assert(typeof taxicabNumbers === 'function');
```

`taxicabNumbers` は配列を返す必要があります。

```js
assert(typeof taxicabNumbers(2) === 'object');
```

`taxicabNumbers` は数値の配列を返す必要があります。

```js
assert(typeof taxicabNumbers(100)[0] === 'number');
```

`taxicabNumbers(4)` は [1729, 4104, 13832, 20683] を返す必要があります。

```js
assert.deepEqual(taxicabNumbers(4), res4);
```

`taxicabNumbers(25)` は [1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656, 110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763, 373464, 402597] を返す必要があります。

```js
assert.deepEqual(taxicabNumbers(25), res25);
```

`taxicabNumbers(39)` の20番目から29番目の結果は [314496,320264,327763,373464,402597,439101,443889,513000,513856] である必要があります。

```js
assert.deepEqual(taxicabNumbers(39).slice(20, 29), res39From20To29);
```

# --seed--

## --seed-contents--

```js
function taxicabNumbers(n) {

  return true;
}
```

# --solutions--

```js
function taxicabNumbers(nNumbers) {
  const cubeN = [];
  const s3s = {};

  const e = 100;
  for (let n = 1; n < e; n += 1) {
    cubeN[n] = n * n * n;
  }

  for (let a = 1; a < e - 1; a += 1) {
    const a3 = cubeN[a];
    for (let b = a; b < e; b += 1) {
      const b3 = cubeN[b];
      const s3 = a3 + b3;

      let abs = s3s[s3];
      if (!abs) {
        s3s[s3] = abs = [];
      }
      abs.push([a, b]);
    }
  }

  let i = 0;
  const res = [];
  Object.keys(s3s).forEach(s3 => {
    const abs = s3s[s3];
    if (abs.length >= 2) { // No two cube pairs found
      i += 1;
      if (i <= nNumbers) {
        res.push(s3);
      }
    }
  });
  return res.map(item => parseInt(item, 10));
}
```
