---
id: 594faaab4e2a8626833e9c3d
title: エスケープ付き文字列のトークン化
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

エスケープされていない区切り文字の出現ごとに文字列を分割できる関数またはプログラムを書いてください。

3つの入力パラメータを受け取る必要があります：

<ul>
  <li><strong>文字列</strong></li>
  <li><strong>区切り文字</strong></li>
  <li><strong>エスケープ文字</strong></li>
</ul>

出力は文字列の配列にしてください。

分割のルール：

<ul>
  <li>区切り文字で分割されたフィールドが出力配列の要素になります。</li>
  <li>空のフィールドも、先頭や末尾にあっても保持してください。</li>
</ul>

エスケープのルール：

<ul>
  <li>「エスケープされている」とは、すでにエスケープされていないエスケープ文字の直後にあることを意味します。</li>
  <li>エスケープ文字が特別な意味を持たない文字の前にあっても、エスケープとしてカウントされます（ただし特別な動作はしません）。</li>
  <li>何かをエスケープするために使われたエスケープ文字は、出力に含めないでください。</li>
</ul>

次のテストケースを満たすことを示してください：

文字列

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

と区切り文字に `|`、エスケープ文字に `^` を使うと、関数は次の配列を出力するはずです：

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

`tokenize` は関数である必要があります。

```js
assert(typeof tokenize === 'function');
```

`tokenize` は配列を返す必要があります。

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` は `['one|uno', '', 'three^^', 'four^|cuatro', '']` を返す必要があります。

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` は `['a&bcd', 'ef', '', '@hi']` を返す必要があります。

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` は `['hello|world', 'howareyou|']` を返す必要があります。

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
