---
id: 5951ed8945deab770972ae56
title: ハノイの塔
challengeType: 1
forumTopicId: 302341
dashedName: towers-of-hanoi
---

# --description--

ハノイの塔の問題を解いてください。最初のパラメータとしてオブジェクトの数が与えられ、その後に各オブジェクトのスタックを識別するための文字列が続きます。移動のリストを含むネストされた配列を作成してください、`["source", "destination"]`。

例えば、パラメータ`(4, 'A', 'B', 'C')`は、移動のネストされた配列`[['A', 'C'], ['B', 'A']]`を返します。これは、1回目の移動がスタック`A`から`C`へオブジェクトを移動し、2回目の移動がスタック`B`から`A`へオブジェクトを移動したことを示しています。

# --instructions--

オブジェクトを積み上げるための移動をネストされた配列で返す関数を書いてください。

# --before-each--

```js
const res3Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B']];
const res7First10Moves = [['A', 'B'], ['A', 'C'], ['B', 'C'], ['A', 'B'], ['C', 'A'], ['C', 'B'], ['A', 'B'], ['A', 'C'], ['B', 'C'], ['B', 'A']];

function getTowerResults_() {
  return {
    res3: towerOfHanoi(3, 'A', 'B', 'C'),
    res5: towerOfHanoi(5, 'X', 'Y', 'Z'),
    res7: towerOfHanoi(7, 'A', 'B', 'C')
  };
}
```

# --hints--

`towerOfHanoi`は関数である必要があります。

```js
assert(typeof towerOfHanoi === 'function');
```

`towerOfHanoi(3, ...)`は7回の移動を返す必要があります。

```js
assert(getTowerResults_().res3.length === 7);
```

`towerOfHanoi(3, 'A', 'B', 'C')`は`[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B']]`を返す必要があります。

```js
assert.deepEqual(getTowerResults_().res3, res3Moves);
```

`towerOfHanoi(5, "X", "Y", "Z")`の10回目の移動はY -> Xであるべきです。

```js
const { res5 } = getTowerResults_();
assert.deepEqual(res5[9], ['Y', 'X']);
```

`towerOfHanoi(7, 'A', 'B', 'C')`の最初の10回の移動は`[['A','B'], ['A','C'], ['B','C'], ['A','B'], ['C','A'], ['C','B'], ['A','B'], ['A','C'], ['B','C'], ['B','A']]`であるべきです。

```js
const { res7 } = getTowerResults_();
assert.deepEqual(res7.slice(0, 10), res7First10Moves);
```

# --seed--

## --seed-contents--

```js
function towerOfHanoi(n, a, b, c) {

  return [[]];
}
```

# --solutions--

```js
function towerOfHanoi(n, a, b, c) {
  const res = [];
  towerOfHanoiHelper(n, a, c, b, res);
  return res;
}

function towerOfHanoiHelper(n, a, b, c, res) {
  if (n > 0) {
    towerOfHanoiHelper(n - 1, a, c, b, res);
    res.push([a, c]);
    towerOfHanoiHelper(n - 1, b, a, c, res);
  }
}
```
