---
id: 595608ff8bcd7a50bd490181
title: ヘイルストーンシーケンス
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

ヘイルストーンシーケンスは、正の整数 `n` から次のように生成できます：

- `n` が `1` の場合、シーケンスは終了します
- `n` が `even` の場合、シーケンスの次の `n` は `= n/2` です
- `n` が `odd` の場合、シーケンスの次の `n` は `= (3 * n) + 1` です

（未証明の）コラッツ予想は、任意の開始数に対してヘイルストーンシーケンスが必ず終了するというものです。

ヘイルストーンシーケンスは、通常、雲の中の雹のように値が何度も下降と上昇を繰り返すため、ヘイルストーン数とも呼ばれ、またコラッツシーケンスとも呼ばれます。

# --instructions--

1. 数字のヘイルストーンシーケンスを生成するルーチンを作成してください
2. 関数は、`limit` 未満の数のうち最も長いヘイルストーンシーケンスを持つ数と、そのシーケンスの長さを配列で返してください。（ただし実際のシーケンスは表示しないでください！）

# --hints--

`hailstoneSequence` は関数であるべきです。

```js
assert(typeof hailstoneSequence === 'function');
```

`hailstoneSequence(30)` は配列を返すべきです。

```js
assert(Array.isArray(hailstoneSequence(30)));
```

`hailstoneSequence(30)` は `[27, 112]` を返すべきです。

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

`hailstoneSequence(50000)` は `[35655, 324]` を返すべきです。

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

`hailstoneSequence(100000)` は `[77031, 351]` を返すべきです。

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
