---
id: 595668ca4cfe1af2fb9818d4
title: ハーシャッド数またはニヴェン数のシーケンス
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

ハーシャッド数またはニヴェン数は、各桁の数字の合計で割り切れる正の整数（1以上）です。

例えば、`42` はハーシャッド数です。なぜなら `42` は `(4 + 2)` で割り切れるからです。

このシーケンスは、増加順に並んだ数として定義されていると仮定します。

# --instructions--

ハーシャッド数のシーケンスの連続するメンバーを生成する関数を実装してください。

それを使って、`n` より大きい最初のハーシャッド数から始まるシーケンスのメンバー10個を含む配列を返してください。

# --hints--

`isHarshadOrNiven` は関数であるべきです。

```js
assert(typeof isHarshadOrNiven === 'function');
```

`isHarshadOrNiven(10)` は `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]` を返すべきです。

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

`isHarshadOrNiven(400)` は `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]` を返すべきです。

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

`isHarshadOrNiven(1000)` は `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]` を返すべきです。

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
