---
id: 59637c4d89f6786115efd814
title: ホフスタッターQシーケンス
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

ホフスタッターQシーケンスは次のように定義されます：

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\\big(n-Q(n-1)\\big)+Q\\big(n-Q(n-2)), \\quad n>2.$

フィボナッチ数列のように定義されていますが、フィボナッチ数列では次の項は前の2つの項の和ですが、Qシーケンスでは前の2つの項が、次の項を作るために足すべき2つの数をQシーケンスのどこまで遡るかを示しています。

# --instructions--

ホフスタッターQシーケンスの方程式を関数として実装してください。その関数は数値`n`を受け取り、整数を返す必要があります。

# --before-each--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

# --hints--

`hofstadterQ`は関数であるべきです。

```js
assert(typeof hofstadterQ === 'function');
```

`hofstadterQ()`は`integer`を返すべきです。

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

`hofstadterQ(1000)`は`502`を返すべきです。

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

`hofstadterQ(1500)`は`755`を返すべきです。

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

`hofstadterQ(2000)`は`1005`を返すべきです。

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

`hofstadterQ(2500)`は`1261`を返すべきです。

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
