---
id: 59669d08d75b60482359409f
title: 日付フォーマット
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

現在の日付を次の仕様で表す2つの文字列を含む配列を返してください。

- 1つ目の文字列は、年、月、日をダッシュ（`-`）で区切った順序で表します。
- 1つ目の文字列の年は4桁で表します。
- 1つ目の文字列の月と日は先頭にゼロを付けません。
- 2つ目の文字列の曜日名と月名は省略せずに表します。
- 2つ目の文字列の日は先頭にゼロを付けません。

出力例：

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

`getDateFormats` は関数である必要があります。

```js
assert(typeof getDateFormats === 'function');
```

`getDateFormats` はオブジェクトを返す必要があります。

```js
assert(typeof getDateFormats() === 'object');
```

`getDateFormats` は2つの要素を持つ配列を返す必要があります。

```js
assert(getDateFormats().length === 2);
```

`getDateFormats` は正しい形式で正しい日付を返す必要があります。

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
