---
id: 5966f99c45e8976909a85575
title: 曜日
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

ある会社では、クリスマスが日曜日にあたる場合、その週の12月25日から1月1日までの間に祝日があれば、従業員に追加の有給休暇を与えて翌週は働かなくてよいようにしています。

# --instructions--

開始年と終了年を受け取り、12月25日が日曜日になるすべての年を配列で返す関数を書いてください。

# --before-each--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

# --hints--

`findXmasSunday` は関数であるべきです。

```js
assert(typeof findXmasSunday === 'function');
```

`findXmasSunday(2000, 2100)` は配列を返すべきです。

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

`findXmasSunday(1970, 2017)` は `[1977, 1983, 1988, 1994, 2005, 2011, 2016]` を返すべきです。

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

`findXmasSunday(2008, 2121)` は `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]` を返すべきです。

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
