---
id: 59667989bf71cf555dd5d2ff
title: S式
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">S式</a>はデータを解析して格納する便利な方法の一つです。

# --instructions--

引用符付きおよび引用符なしの文字列、整数、浮動小数点数を扱うシンプルなS式リーダー／パーサーを書いてください。

関数は文字列から単一のネストされたS式を読み取り、それを（ネストされた）配列として返す必要があります。

改行やその他の空白は、引用符付き文字列内に含まれていない限り無視されます。

引用符付き文字列内の "`()`" は解釈されず、文字列の一部として扱われます。

文字列内のエスケープされた引用符の処理は任意です。したがって "`(foo"bar)`" は文字列 "`foo"bar`" として扱ってもエラーとして扱ってもかまいません。

このため、リーダーはエスケープのための `\` を認識する必要はありませんが、言語に適切なデータ型があれば数値は認識する必要があります。

`()"`（エスケープがサポートされていれば `\`）と空白を除き、特殊文字はありません。その他のものは引用符なしで許可されます。

リーダーは以下の入力を読み取ることができる必要があります。

<pre>((data "quoted data" 123 4.5)
(data (!@# (4.5) "(more" "data)")))
</pre>

そしてそれをネイティブなデータ構造に変換してください。

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

`parseSexpr` は関数である必要があります。

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` は `['data1', 'data2', 'data3']` を返す必要があります。

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` は `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]` を返す必要があります。

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
