---
id: 59694356a6e7011f7f1c5f4e
title: FreeCellのカードを配る
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* は、1978年にPaul AlfilleがPLATOシステムに導入したソリティアカードゲームです。MicrosoftのJim Horneが名前をFreeCellに変更し、DOSおよびWindows向けにゲームを再実装しました。このバージョンでは32000の番号付きディールが導入されました。

ゲームが人気になるにつれて、Jim Horneはアルゴリズムを公開し、他のFreeCellの実装もMicrosoftのディールを再現し始めました。これらのディールは1から32000まで番号が付けられています。Microsoftの新しいバージョンでは100万のディールがあり、1から1000000まで番号が付けられています。一部の実装ではその範囲外の番号も許可されています。

アルゴリズムはMicrosoft Cのこの線形合同生成器を使用します：

<ul>
  <li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$ は0から32767の範囲です。</li>
</ul>

アルゴリズムは次の通りです：

<ol>
  <li>ディール番号を使って乱数生成器を初期化します。
  </li><li>52枚のカードの配列を作成します：クラブのエース、ダイヤのエース、ハートのエース、スペードのエース、クラブの2、ダイヤの2、…とランク順に：エース、2、3、4、5、6、7、8、9、10、ジャック、クイーン、キング。配列のインデックスは0から51で、クラブのエースが0、スペードのキングが51です。</li>
  <li>配列が空になるまで：</li>
    <ul>
      <li>ランダムなカードを<i>インデックス</i> ≡ <i>次の乱数</i> (mod <i>配列の長さ</i>)で選びます。</li>
      <li>このランダムなカードと配列の最後のカードを入れ替えます。</li>
      <li>このランダムなカードを配列から取り除きます。（配列の長さは1減ります。）</li>
      <li>このランダムなカードを配ります。</li>
    </ul>
  <li>52枚すべてのカードを表向きで8列に配ります。最初の8枚は8列に1枚ずつ、次の8枚はそれぞれの列の上に、というように続けます。</li>
</ol>

**例：**

**カードを配る順番**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**ゲーム #1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**ゲーム #617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

ディール番号を受け取り、このアルゴリズムと同じ順番でカードを配る関数を書いてください。関数はFreeCellのボードを表す二次元配列を返す必要があります。

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` は関数であるべきです。

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` はオブジェクトを返すべきです。

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` は長さ7の配列を返すべきです。

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` は例の「ゲーム #1」と同じ配列を返すべきです。

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` は例の「ゲーム #617」と同じ配列を返すべきです。

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
