---
id: 596a8888ab7c01048de257d5
title: ディープコピー
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

与えられたオブジェクトのディープコピーを返す関数を書いてください。コピーは元のオブジェクトと同じものではいけません。

この課題では以下のケースはテストされません：

<ul>
  <li>プロパティが関数であるオブジェクト</li>
  <li>日付オブジェクトやプロパティに日付オブジェクトを持つオブジェクト</li>
  <li>正規表現オブジェクトやプロパティに正規表現オブジェクトを持つオブジェクト</li>
  <li>プロトタイプのコピー</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

`deepcopy` は関数である必要があります。

```js
assert(typeof deepcopy === 'function');
```

`deepcopy({test: "test"})` はオブジェクトを返す必要があります。

```js
assert(typeof deepcopy(obj1) === 'object');
```

`deepcopy` は渡されたオブジェクトと同じものを返してはいけません。

```js
assert(deepcopy(obj2) != obj2);
```

配列を含むオブジェクトが渡された場合、`deepcopy` はそのオブジェクトのディープコピーを返す必要があります。

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

別のオブジェクトを含むオブジェクトが渡された場合、`deepcopy` はそのオブジェクトのディープコピーを返す必要があります。

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
