---
id: 596e457071c35c882915b3e4
title: 文字列のリストを比較する
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

リストは順序付けられた値の集合で、重複を含む場合があります。例を示します：

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

任意の数の文字列のリストが与えられたとき、次の条件それぞれに対して関数を実装してください：

<ul>
  <li>すべての文字列が辞書順で等しいかどうかをテストする</li>
  <li>各文字列が次の文字列より辞書順で小さいかどうかをテストする（つまり、リストが厳密な昇順になっているかどうか）</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

`allEqual` は関数である必要があります。

```js
assert(typeof allEqual === 'function');
```

`azSorted` は関数である必要があります。

```js
assert(typeof azSorted === 'function');
```

`allEqual(["AA", "AA", "AA", "AA"])` は true を返す必要があります。

```js
assert(allEqual(testCases[0]));
```

`azSorted(["AA", "AA", "AA", "AA"])` は false を返す必要があります。

```js
assert(!azSorted(testCases[0]));
```

`allEqual(["AA", "ACB", "BB", "CC"])` は false を返す必要があります。

```js
assert(!allEqual(testCases[1]));
```

`azSorted(["AA", "ACB", "BB", "CC"])` は true を返す必要があります。

```js
assert(azSorted(testCases[1]));
```

`allEqual([])` は true を返す必要があります。

```js
assert(allEqual(testCases[2]));
```

`azSorted([])` は true を返す必要があります。

```js
assert(azSorted(testCases[2]));
```

`allEqual(["AA"])` は true を返す必要があります。

```js
assert(allEqual(testCases[3]));
```

`azSorted(["AA"])` は true を返す必要があります。

```js
assert(azSorted(testCases[3]));
```

`allEqual(["BB", "AA"])` は false を返す必要があります。

```js
assert(!allEqual(testCases[4]));
```

`azSorted(["BB", "AA"])` は false を返す必要があります。

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
