---
id: 596e414344c3b2872167f0fe
title: カンマの言い争い
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">カンマの言い争い</a>は、Eric Lippertが自身のブログで最初に出題した課題です。

# --instructions--

リストやシーケンスからの単語の連結で文字列を生成する関数を書いてください。条件は以下の通りです：

<ol>
  <li>単語が1つもない入力の場合、出力は中括弧2つだけの文字列（<code>"{}"</code>）になります。</li>
  <li>単語が1つだけの場合、例えば<code>["ABC"]</code>、出力はその単語を中括弧で囲んだ文字列、例えば<code>"{ABC}"</code>になります。</li>
  <li>単語が2つの場合、例えば<code>["ABC", "DEF"]</code>、出力は中括弧で囲み、単語の間を<code>" and "</code>でつなげた文字列、例えば<code>"{ABC and DEF}"</code>になります。</li>
  <li>単語が3つ以上の場合、例えば<code>["ABC", "DEF", "G", "H"]</code>、最後の単語以外は<code>", "</code>で区切り、最後の単語は<code>" and "</code>でつなぎ、すべて中括弧で囲んだ文字列、例えば<code>"{ABC, DEF, G and H}"</code>になります。</li>
</ol>

以下の入力例で関数をテストし、ここに出力を表示してください：

<ul>
  <li>[] # （単語なし）</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**注意：** この課題では単語はすべて大文字の空でない文字列と仮定してください。

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` は関数である必要があります。

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` は文字列を返す必要があります。

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` は "{}" を返す必要があります。

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` は `"{ABC}"` を返す必要があります。

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` は `"{ABC and DEF}"` を返す必要があります。

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` は `"{ABC, DEF, G and H}"` を返す必要があります。

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
