---
id: 596fda99c69f779975a1b67d
title: 部分文字列の出現回数を数える
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

文字列の中にある部分文字列の重ならない出現回数を数える関数、または組み込み関数を示してください。

関数は2つの引数を取る必要があります。

<ul>
  <li>最初の引数は検索対象の文字列、</li>
  <li>2つ目は検索する部分文字列です。</li>
</ul>

整数のカウントを返してください。

マッチングは重ならない最大数の一致を返す必要があります。

一般的には、これは左から右、または右から左へのマッチングを意味します。

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

`countSubstring` は関数であるべきです。

```js
assert(typeof countSubstring === 'function');
```

`countSubstring("the three truths", "th")` は `3` を返すべきです。

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

`countSubstring("ababababab", "abab")` は `2` を返すべきです。

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

`countSubstring("abaabba*bbaba*bbab", "a*b")` は `2` を返すべきです。

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
