---
id: 596fd036dc1ab896c5db98b1
title: 秒数を複合期間に変換する
challengeType: 1
forumTopicId: 302236
dashedName: convert-seconds-to-compound-duration
---

# --description--

次のことを行う関数を実装してください。

<ul>
  <li>秒数で表された正の整数の期間を入力として受け取り（例：<code>100</code>）、</li>
  <li>同じ期間を週、日、時間、分、秒に分解して示す文字列を返します（例：<code>1 min, 40 sec</code>）。</li>
</ul>

以下の3つのテストケースを通過することを示してください。

<div style='font-size:115%; font-weight: bold;'>テストケース</div>

| 入力の数値 | 出力の文字列                         |
| ------------ | ------------------------------------- |
| 7259         | <code>2 hr, 59 sec</code>             |
| 86400        | <code>1 d</code>                      |
| 6000000      | <code>9 wk, 6 d, 10 hr, 40 min</code> |

<div style="font-size:115%; font-weight: bold;">詳細</div>
<ul>
  <li>
    以下の5つの単位を使用してください。

| 単位   | 出力で使う接尾辞       | 変換関係              |
| ------ | --------------------- | --------------------- |
| 週     | <code>wk</code>       | 1週 = 7日             |
| 日     | <code>d</code>        | 1日 = 24時間          |
| 時間   | <code>hr</code>       | 1時間 = 60分          |
| 分     | <code>min</code>      | 1分 = 60秒            |
| 秒     | <code>sec</code>      | ---                   |

  </li>
  <li>
    ただし、出力には<strong>ゼロでない値の数量のみ</strong>含めてください（例：<code>1 d</code>を返し、<code>0 wk, 1 d, 0 hr, 0 min, 0 sec</code>は返さないでください）。
  </li>
  <li>
    可能な限り大きい単位を優先してください（例：<code>2 min, 10 sec</code>を返し、<code>1 min, 70 sec</code>や<code>130 sec</code>は返さないでください）。
  </li>
  <li>
    テストケースの書式を真似てください（数量は大きい単位から小さい単位の順に並べ、カンマ＋スペースで区切り、各数量の値と単位はスペースで区切る）。
  </li>
</ul>

# --before-each--

```js
const testCases = [7259, 86400, 6000000];
const results = ['2 hr, 59 sec', '1 d', '9 wk, 6 d, 10 hr, 40 min'];
```

# --hints--

`convertSeconds` は関数であるべきです。

```js
assert(typeof convertSeconds === 'function');
```

`convertSeconds(7259)` は `2 hr, 59 sec` を返すべきです。

```js
assert.equal(convertSeconds(testCases[0]), results[0]);
```

`convertSeconds(86400)` は `1 d` を返すべきです。

```js
assert.equal(convertSeconds(testCases[1]), results[1]);
```

`convertSeconds(6000000)` は `9 wk, 6 d, 10 hr, 40 min` を返すべきです。

```js
assert.equal(convertSeconds(testCases[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function convertSeconds(sec) {

  return true;
}
```

# --solutions--

```js
function convertSeconds(sec) {
  const localNames = ['wk', 'd', 'hr', 'min', 'sec'];
  // compoundDuration :: [String] -> Int -> String
  const compoundDuration = (labels, intSeconds) =>
    weekParts(intSeconds)
    .map((v, i) => [v, labels[i]])
    .reduce((a, x) =>
      a.concat(x[0] ? [`${x[0]} ${x[1] || '?'}`] : []), []
    )
    .join(', ');

    // weekParts :: Int -> [Int]
  const weekParts = intSeconds => [0, 7, 24, 60, 60]
    .reduceRight((a, x) => {
      const r = a.rem;
      const mod = x !== 0 ? r % x : r;

      return {
        rem: (r - mod) / (x || 1),
        parts: [mod].concat(a.parts)
      };
    }, {
      rem: intSeconds,
      parts: []
    })
    .parts;

  return compoundDuration(localNames, sec);
}
```
