---
id: 597089c87eec450c68aa1643
title: プリミティブデータ型を定義する
challengeType: 1
forumTopicId: 302248
dashedName: define-a-primitive-data-type
---

# --description--

整数のように振る舞いますが、有効な最小値が1、最大値が10の型を定義してください。

エラー処理：

<ul>
  <li><code>Num</code>を1から10の範囲外の値でインスタンス化しようとすると、エラーメッセージが<code>'Out of range'</code>の<code>TypeError</code>をスローする必要があります。</li>
  <li><code>Num</code>を数値以外の値でインスタンス化しようとすると、エラーメッセージが<code>'Not a Number'</code>の<code>TypeError</code>をスローする必要があります。</li>
</ul>

# --hints--

`Num` は関数である必要があります。

```js
assert(typeof Num === 'function');
```

`new Num(4)` はオブジェクトを返す必要があります。

```js
assert(typeof new Num(4) === 'object');
```

`new Num('test')` はメッセージが 'Not a Number' の TypeError をスローする必要があります。

```js
assert.throws(() => new Num('test'), TypeError);
```

`new Num(0)` はメッセージが 'Out of range' の TypeError をスローする必要があります。

```js
assert.throws(() => new Num(0), TypeError);
```

`new Num(-5)` はメッセージが 'Out of range' の TypeError をスローする必要があります。

```js
assert.throws(() => new Num(-5), TypeError);
```

`new Num(10)` はメッセージが 'Out of range' の TypeError をスローする必要があります。

```js
assert.throws(() => new Num(11), TypeError);
```

`new Num(20)` はメッセージが 'Out of range' の TypeError をスローする必要があります。

```js
assert.throws(() => new Num(20), TypeError);
```

`new Num(3) + new Num(4)` は7と等しくなる必要があります。

```js
assert.equal(new Num(3) + new Num(4), 7);
```

`new Num(3) - new Num(4)` は-1と等しくなる必要があります。

```js
assert.equal(new Num(3) - new Num(4), -1);
```

`new Num(3) * new Num(4)` は12と等しくなる必要があります。

```js
assert.equal(new Num(3) * new Num(4), 12);
```

`new Num(3) / new Num(4)` は0.75と等しくなる必要があります。

```js
assert.equal(new Num(3) / new Num(4), 0.75);
```

`new Num(3) < new Num(4)` はtrueである必要があります。

```js
assert(new Num(3) < new Num(4));
```

`new Num(3) > new Num(4)` はfalseである必要があります。

```js
assert(!(new Num(3) > new Num(4)));
```

`(new Num(5)).toString()` は '5' を返す必要があります。

```js
assert.equal(new Num(5).toString(), '5');
```

# --seed--

## --seed-contents--

```js
function Num(n) {

  return n;
}
```

# --solutions--

```js
function Num(n) {
  if (isNaN(n)) {
    throw new TypeError('Not a Number');
  }
  if (n < 1 || n > 10) {
    throw new TypeError('Out of range');
  }

  this._value = +n;
}
Num.prototype.valueOf = function() { return this._value; };
Num.prototype.toString = function () { return this._value.toString(); };
```
