---
id: 59713bd26bdeb8a594fb9413
title: コインの数え方
challengeType: 1
forumTopicId: 302238
dashedName: count-the-coins
---

# --description--

アメリカの通貨には4種類の一般的な硬貨があります：

<ul>
  <li>クォーター（25セント）</li>
  <li>ダイム（10セント）</li>
  <li>ニッケル（5セント）、および</li>
  <li>ペニー（1セント）</li>
</ul>

<p>15セントを作る方法は6通りあります：</p>

<ul>
  <li>ダイム1枚とニッケル1枚</li>
  <li>ダイム1枚とペニー5枚</li>
  <li>ニッケル3枚</li>
  <li>ニッケル2枚とペニー5枚</li>
  <li>ニッケル1枚とペニー10枚</li>
  <li>ペニー15枚</li>
</ul>

# --instructions--

アメリカの一般的な硬貨を使って、`cents`で表されるアメリカのペニーの金額に対して、いくつの組み合わせでおつりを作れるかを判定する関数を実装してください。

# --hints--

`countCoins`は関数である必要があります。

```js
assert(typeof countCoins === 'function');
```

`countCoins(15)`は`6`を返す必要があります。

```js
assert.equal(countCoins(15), 6);
```

`countCoins(85)`は`163`を返す必要があります。

```js
assert.equal(countCoins(85), 163);
```

`countCoins(100)`は`242`を返す必要があります。

```js
assert.equal(countCoins(100), 242);
```

# --seed--

## --seed-contents--

```js
function countCoins(cents) {

  return true;
}
```

# --solutions--

```js
function countCoins(cents) {
  const operands = [1, 5, 10, 25];
  const targetsLength = cents + 1;
  const operandsLength = operands.length;
  const t = [1];

  for (let a = 0; a < operandsLength; a++) {
    for (let b = 1; b < targetsLength; b++) {
      // initialise undefined target
      t[b] = t[b] ? t[b] : 0;

      // accumulate target + operand ways
      t[b] += (b < operands[a]) ? 0 : t[b - operands[a]];
    }
  }

  return t[targetsLength - 1];
}
```
