---
id: 597f24c1dda4e70f53c79c81
title: フィボナッチ数列
challengeType: 1
forumTopicId: 302268
dashedName: fibonacci-sequence
---

# --description--

<code>n<sup>番目</sup></code>のフィボナッチ数を生成する関数を書いてください。

<code>n<sup>番目</sup></code>のフィボナッチ数は次のように表されます：

<code>F<sub>n</sub> = F<sub>n-1</sub> + F<sub>n-2</sub></code>

この数列の最初の2つの項は0と1です。

したがって、数列は次のようになります：0、1、1、2、3、5、8、13... 

# --hints--

`fibonacci` は関数である必要があります。

```js
assert(typeof fibonacci === 'function');
```

`fibonacci(2)` は数値を返す必要があります。

```js
assert(typeof fibonacci(2) == 'number');
```

`fibonacci(3)` は2を返す必要があります。

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` は5を返す必要があります。

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` は55を返す必要があります。

```js
assert.equal(fibonacci(10), 55);
```

# --seed--

## --seed-contents--

```js
function fibonacci(n) {

}
```

# --solutions--

```js
function fibonacci(n) {
  let a = 0, b = 1, t;
  while (--n >= 0) {
    t = a;
    a = b;
    b += t;
  }
  return a;
}
```
