---
id: 5987fd532b954e0f21b5d3f6
title: 平衡インデックス
challengeType: 1
forumTopicId: 302255
dashedName: equilibrium-index
---

# --description--

シーケンスの平衡インデックスとは、低いインデックスの要素の合計が高いインデックスの要素の合計と等しくなるようなシーケンスのインデックスのことです。

例えば、シーケンス $A$ において：

<ul style='list-style: none;'>
  <li><big>$A_0 = -7$</big></li>
  <li><big>$A_1 =  1$</big></li>
  <li><big>$A_2 =  5$</big></li>
  <li><big>$A_3 =  2$</big></li>
  <li><big>$A_4 = -4$</big></li>
  <li><big>$A_5 =  3$</big></li>
  <li><big>$A_6 =  0$</big></li>
</ul>

`3` は平衡インデックスです。なぜなら：

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 = A_4 + A_5 + A_6$</big></li>
</ul>

`6` も平衡インデックスです。なぜなら：

<ul style='list-style: none;'>
  <li><big>$A_0 + A_1 + A_2 + A_3 + A_4 + A_5 = 0$</big></li>
</ul>

（要素がゼロ個の合計はゼロです）

`7` は平衡インデックスではありません。なぜならシーケンス $A$ の有効なインデックスではないからです。

# --instructions--

シーケンスが与えられたとき、その平衡インデックス（存在する場合）を返す関数を書いてください。

シーケンスは非常に長い場合があると想定してください。

# --hints--

`equilibrium` は関数であるべきです。

```js
assert(typeof equilibrium === 'function');
```

`equilibrium([-7, 1, 5, 2, -4, 3, 0])` は `[3,6]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[0]),
  globalThis.ans[0]
);
```

`equilibrium([2, 4, 6])` は `[]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[1]),
  globalThis.ans[1]
);
```

`equilibrium([2, 9, 2])` は `[1]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[2]),
  globalThis.ans[2]
);
```

`equilibrium([1, -1, 1, -1, 1, -1, 1])` は `[0,1,2,3,4,5,6]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[3]),
  globalThis.ans[3]
);
```

`equilibrium([1])` は `[0]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[4]),
  globalThis.ans[4]
);
```

`equilibrium([])` は `[]` を返すべきです。

```js
assert.deepEqual(
  equilibrium(globalThis.equilibriumTests[5]),
  globalThis.ans[5]
);
```

# --before-all--

```js
globalThis.equilibriumTests =
  [[-7, 1, 5, 2, -4, 3, 0], // 3, 6
  [2, 4, 6], // empty
  [2, 9, 2], // 1
  [1, -1, 1, -1, 1, -1, 1], // 0,1,2,3,4,5,6
  [1], // 0
  [] // empty
  ];
globalThis.ans = [[3, 6], [], [1], [0, 1, 2, 3, 4, 5, 6], [0], []];
```

# --seed--

## --seed-contents--

```js
function equilibrium(a) {

}
```

# --solutions--

```js
function equilibrium(a) {
  let N = a.length,
    i,
    l = [],
    r = [],
    e = [];
  for (l[0] = a[0], r[N - 1] = a[N - 1], i = 1; i < N; i++)
    { l[i] = l[i - 1] + a[i], r[N - i - 1] = r[N - i] + a[N - i - 1]; }
  for (i = 0; i < N; i++)
    { if (l[i] === r[i]) e.push(i); }
  return e;
}
```
