---
id: 59880443fb36441083c6c20e
title: オイラー法
challengeType: 1
forumTopicId: 302258
dashedName: euler-method
---

# --description--

オイラー法は、初期値が与えられた1階常微分方程式（ODE）の数値解を近似します。これは初期値問題（IVP）を解くための明示的な方法であり、<a href="https://www.freecodecamp.org/news/eulers-method-explained-with-examples/" title="Euler's Method Explained with Examples" target="_blank" rel="noopener noreferrer nofollow">この記事</a>で説明されています。

ODEは次の形で与えられる必要があります：

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt} = f(t,y(t))$</big></li>
</ul>

初期値は

<ul style='list-style: none;'>
  <li><big>$y(t_0) = y_0$</big></li>
</ul>

数値解を得るために、左辺の微分を有限差分近似で置き換えます：

<ul style='list-style: none;'>
  <li><big>$\frac{dy(t)}{dt}  \approx \frac{y(t+h)-y(t)}{h}$</big></li>
</ul>

そして $y(t+h)$ を解きます：

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, \frac{dy(t)}{dt}$</big></li>
</ul>

これは次と同じです：

<ul style='list-style: none;'>
  <li><big>$y(t+h) \approx y(t) + h \, f(t,y(t))$</big></li>
</ul>

反復解のルールは次の通りです：

<ul style='list-style: none;'>
  <li><big>$y_{n+1} = y_n + h \, f(t_n, y_n)$</big></li>
</ul>

ここで $h$ はステップサイズで、解の精度に最も関係するパラメータです。ステップサイズが小さいほど精度は上がりますが計算コストも増えるため、問題に応じて手動で選ぶ必要があります。

**例：ニュートンの冷却法則**

ニュートンの冷却法則は、初期温度 $T(t_0) = T_0$ の物体が温度 $T_R$ の環境でどのように冷えるかを表します：

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, \Delta T$</big></li>
</ul>

または

<ul style='list-style: none;'>
  <li><big>$\frac{dT(t)}{dt} = -k \, (T(t) - T_R)$</big></li>
</ul>

これは、物体の冷却速度 $\\frac{dT(t)}{dt}$ が現在の温度差 $\\Delta T = (T(t) - T_R)$ に比例することを示しています。

解析解は、数値近似と比較するために次のようになります：

<ul style='list-style: none;'>
  <li><big>$T(t) = T_R + (T_0 - T_R) \; e^{-k t}$</big></li>
</ul>

# --instructions--

オイラー法のルーチンを実装し、次の3つの異なるステップサイズでニュートンの冷却法則の例を解いてください：

<ul>
  <li><code>2 s</code></li>
  <li><code>5 s</code></li>
  <li><code>10 s</code></li>
</ul>

そして解析解と比較してください。

**初期値：**

<ul>
  <li>初期温度 <big>$T_0$</big> は <code>100 °C</code></li>
  <li>室温 <big>$T_R$</big> は <code>20 °C</code></li>
  <li>冷却定数 <big>$k$</big> は <code>0.07</code></li>
  <li>計算する時間範囲は <code>0 s</code> から <code>100 s</code></li>
</ul>  

関数の最初のパラメータは初期時間、2番目は初期温度、3番目は経過時間、4番目はステップサイズです。

# --hints--

`eulersMethod` は関数であるべきです。

```js
assert(typeof eulersMethod === 'function');
```

`eulersMethod(0, 100, 100, 2)` は数値を返すべきです。

```js
assert(typeof eulersMethod(0, 100, 100, 2) === 'number');
```

`eulersMethod(0, 100, 100, 2)` は 20.0424631833732 を返すべきです。

```js
assert.equal(eulersMethod(0, 100, 100, 2), 20.0424631833732);
```

`eulersMethod(0, 100, 100, 5)` は 20.01449963666907 を返すべきです。

```js
assert.equal(eulersMethod(0, 100, 100, 5), 20.01449963666907);
```

`eulersMethod(0, 100, 100, 10)` は 20.000472392 を返すべきです。

```js
assert.equal(eulersMethod(0, 100, 100, 10), 20.000472392);
```

# --seed--

## --seed-contents--

```js
function eulersMethod(x1, y1, x2, h) {

}
```

# --solutions--

```js
function eulersMethod(x1, y1, x2, h) {
  let x = x1;
  let y = y1;

  while ((x < x2 && x1 < x2) || (x > x2 && x1 > x2)) {
    y += h * (-0.07 * (y - 20));
    x += h;
  }

  return y;
}
```
