---
id: 598eea87e5cf4b116c3ff81a
title: メルセンヌ数の因数
challengeType: 1
forumTopicId: 302264
dashedName: factors-of-a-mersenne-number
---

# --description--

メルセンヌ数とは、<code>2<sup>P</sup>-1</code> の形の数です。

もし `P` が素数であれば、そのメルセンヌ数はメルセンヌ素数である可能性があります。（`P` が素数でなければ、そのメルセンヌ数も素数ではありません。）

メルセンヌ素数の探索では、潜在的に長時間かかる<a href="https://rosettacode.org/wiki/Lucas-Lehmer test" target="_blank" rel="noopener noreferrer nofollow">Lucas-Lehmerテスト</a>を始める前に、小さな因数を見つけて指数を除外することが有利です。

<code>2<sup>P</sup>-1</code> を割り切るかどうか（または同等に、<code>2<sup>P</sup> mod (その数) = 1</code>かどうか）を判定する非常に効率的なアルゴリズムがあります。

いくつかの言語には、この指数とモジュラスの演算（modPowなどと呼ばれる）の組み込み実装があります。

以下は、この modPow を自分で実装する方法です。

例えば、<code>2<sup>23</sup> mod 47</code> を計算してみましょう。

指数23を2進数に変換すると10111になります。<code><tt>square</tt> = 1</code> から始めて、繰り返し自乗します。

指数の最上位ビットを取り除き、それが1ならば `square` に指数の底（2）を乗算し、次に <code><tt>square</tt> を47で割った余りを計算します。

前のステップのモジュラスの結果を次のステップの `square` の初期値として使います：

<pre>取り除く   オプション
square        最上位ビット  2を乗算  47で割った余り
------------  -------  -------------  ------
1*1 = 1       1  0111  1*2 = 2           2
2*2 = 4       0   111     なし           4
4*4 = 16      1    11  16*2 = 32        32
32*32 = 1024  1     1  1024*2 = 2048    27
27*27 = 729   1        729*2 = 1458      1
</pre>

<code>2<sup>23</sup> mod 47 = 1</code> なので、47は <code>2<sup>P</sup>-1</code> の因数です。

（これを見るには、両辺から1を引いてください：<code>2<sup>23</sup>-1 = 0 mod 47</code>。）

47が因数であることが示されたので、<code>2<sup>23</sup>-1</code> は素数ではありません。

メルセンヌ数のさらなる性質により、この処理をさらに洗練できます。

<code>2<sup>P</sup>-1</code> の因数 `q` は、`2kP+1` の形でなければならず、`k` は正の整数または0です。さらに、`q` は `1` または `7 mod 8` でなければなりません。

最後に、潜在的な因数 `q` は<a href="https://rosettacode.org/wiki/Primality_by_trial_division" target="_blank" rel="noopener noreferrer nofollow">素数</a>である必要があります。

他の試し割りアルゴリズムと同様に、アルゴリズムは `2kP+1 > sqrt(N)` のときに停止します。これらの主なテストは `P` が素数であるメルセンヌ数にのみ有効です。例えば、<code>M<sub>4</sub>=15</code> はこれらの手法では因数が見つかりませんが、3と5に因数分解され、どちらも `2kP+1` に合致しません。

# --instructions--

上記の方法を使って、<code>2<sup>p</sup>-1</code> の因数を見つけてください。

# --hints--

`check_mersenne` は関数であるべきです。

```js
assert(typeof check_mersenne === 'function');
```

`check_mersenne(3)` は文字列を返すべきです。

```js
assert(typeof check_mersenne(3) == 'string');
```

`check_mersenne(3)` は文字列 `M3 = 2^3-1 is prime` を返すべきです。

```js
assert.equal(check_mersenne(3), 'M3 = 2^3-1 is prime');
```

`check_mersenne(23)` は文字列 `M23 = 2^23-1 is composite with factor 47` を返すべきです。

```js
assert.equal(check_mersenne(23), 'M23 = 2^23-1 is composite with factor 47');
```

`check_mersenne(929)` は文字列 `M929 = 2^929-1 is composite with factor 13007` を返すべきです。

```js
assert.equal(
  check_mersenne(929),
  'M929 = 2^929-1 is composite with factor 13007'
);
```

# --seed--

## --seed-contents--

```js
function check_mersenne(p) {

}
```

# --solutions--

```js
function check_mersenne(p){
    function isPrime(value){
      for (let i=2; i < value; i++){
        if (value % i == 0){
          return false;
        }
        if (value % i != 0){
          return true;
         }
      }
    }

    function trial_factor(base, exp, mod){
      let square, bits;
      square = 1;
      bits = exp.toString(2).split('');
      for (let i=0,ln=bits.length; i<ln; i++){
        square = Math.pow(square, 2) * (bits[i] == 1 ? base : 1) % mod;
      }
      return (square == 1);
    }

    function mersenne_factor(p){
      let limit, k, q;
      limit = Math.sqrt(Math.pow(2,p) - 1);
      k = 1;
      while ((2*k*p - 1) < limit){
        q = 2*k*p + 1;
        if (isPrime(q) && (q % 8 == 1 || q % 8 == 7) && trial_factor(2,p,q)){
          return q; // q is a factor of 2**p-1
        }
        k++;
      }
      return null;
    }
  let f, result;
  result="M"+p+" = 2^"+p+"-1 is ";
  f = mersenne_factor(p);
  result+=f == null ? "prime" : "composite with factor "+f;
  return result;
}
```
