---
id: 598ee8b91b410510ae82efef
title: 拡張可能な素数ジェネレーター
challengeType: 1
forumTopicId: 302262
dashedName: extensible-prime-generator
---

# --description--

順番に素数を生成するジェネレーターを作成してください。任意の十分に大きな素数の生成にも自動で対応できるようにしてください。

ジェネレーターは以下のことができる必要があります：

<ul>
  <li>最初の<code>n</code>個の素数を表示する</li>
  <li>範囲内の素数を表示する</li>
  <li>範囲内の素数の数を表示する</li>
  <li><code>n<sup>番目</sup></code>の素数を表示する</li>
</ul>

関数は2つのパラメータを持ちます。1つ目は`n`または範囲を配列として受け取ります。2つ目はブール値で、関数が素数を配列として返すか、単一の数値（範囲内の素数の数または<code>n<sup>番目</sup></code>の素数）を返すかを指定します。パラメータに応じて関数は配列を返す必要があります。

# --hints--

`primeGenerator`は関数である必要があります。

```js
assert(typeof primeGenerator === 'function');
```

`primeGenerator(20, true)`は`[2, 3, 5, 7, 11, 13, 17, 19, 23, 29, 31, 37, 41, 43, 47, 53, 59, 61, 67, 71]`を返す必要があります。

```js
assert.deepEqual(primeGenerator(20, true), [
  2,
  3,
  5,
  7,
  11,
  13,
  17,
  19,
  23,
  29,
  31,
  37,
  41,
  43,
  47,
  53,
  59,
  61,
  67,
  71
]);
```

`primeGenerator([100, 150], true)`は`[101, 103, 107, 109, 113, 127, 131, 137, 139, 149]`を返す必要があります。

```js
assert.deepEqual(primeGenerator([100, 150], true), [
  101,
  103,
  107,
  109,
  113,
  127,
  131,
  137,
  139,
  149
]);
```

`primeGenerator([7700, 8000], false)`は`30`を返す必要があります。

```js
assert.equal(primeGenerator([7700, 8000], false), 30);
```

`primeGenerator(10000, false)`は`104729`を返す必要があります。

```js
assert.equal(primeGenerator(10000, false), 104729);
```

# --seed--

## --seed-contents--

```js
function primeGenerator(num, showPrimes) {

}
```

# --solutions--

```js
function primeGenerator(num, showPrimes) {
  let i,
    arr = [];

  function isPrime(num) {
    // try primes <= 16
    if (num <= 16) { return (
      num == 2 || num == 3 || num == 5 || num == 7 || num == 11 || num == 13
    ); }
    // cull multiples of 2, 3, 5 or 7
    if (num % 2 == 0 || num % 3 == 0 || num % 5 == 0 || num % 7 == 0)
      { return false; }
    // cull square numbers ending in 1, 3, 7 or 9
    for (let i = 10; i * i <= num; i += 10) {
      if (num % (i + 1) == 0) return false;
      if (num % (i + 3) == 0) return false;
      if (num % (i + 7) == 0) return false;
      if (num % (i + 9) == 0) return false;
    }
    return true;
  }

  if (typeof num === 'number') {
    for (i = 0; arr.length < num; i++) if (isPrime(i)) arr.push(i);
    // first x primes
    if (showPrimes) return arr;
    // xth prime
    return arr.pop();
  }

  if (Array.isArray(num)) {
    for (i = num[0]; i <= num[1]; i++) if (isPrime(i)) arr.push(i);
    // primes between x .. y
    if (showPrimes) return arr;
    // number of primes between x .. y
    return arr.length;
  }
}
```
