---
id: 598eef80ba501f1268170e1e
title: フィボナッチ n ステップ数列
challengeType: 1
forumTopicId: 302267
dashedName: fibonacci-n-step-number-sequences
---

# --description--

これらの数列は、通常のフィボナッチ数列を拡張したものです：

<ol>
  <li>$n = 2$ の場合はフィボナッチ数列で、初期値は $[1, 1]$、$F_k^2 = F_{k-1}^2 + F_{k-2}^2$ となります</li>
  <li>$n = 3$ の場合はトリボナッチ数列で、初期値は $[1, 1, 2]$、$F_k^3 = F_{k-1}^3 + F_{k-2}^3 + F_{k-3}^3$ となります</li>
  <li>$n = 4$ の場合はテトラナッチ数列で、初期値は $[1, 1, 2, 4]$、$F_k^4 = F_{k-1}^4 + F_{k-2}^4 + F_{k-3}^4 + F_{k-4}^4$ となります…</li>
  <li>一般の $n>2$ の場合はフィボナッチ $n$ ステップ数列 $F_k^n$ で、初期値は $(n-1)$ 番目のフィボナッチ $n$ ステップ数列 $F_k^{n-1}$ の最初の $n$ 個の値で、$n$ 番目の数列の $k$ 番目の値は $F_k^n = \sum_{i=1}^{(n)} {F_{k-i}^{(n)}}$ となります</li>
</ol>

小さい $n$ の値の場合、ギリシャ語の数詞接頭辞が使われてそれぞれの数列に名前が付けられることがあります。

フィボナッチ $n$ ステップ数列：

| $n$ | 数列名       | 値の例                                                  |
| --- | ------------ | ------------------------------------------------------- |
| 2   | fibonacci    | 1 1 2 3 5 8 13 21 34 55 89 144 233 377 610 ...          |
| 3   | tribonacci   | 1 1 2 4 7 13 24 44 81 149 274 504 927 1705 3136 ...     |
| 4   | tetranacci   | 1 1 2 4 8 15 29 56 108 208 401 773 1490 2872 5536 ...   |
| 5   | pentanacci   | 1 1 2 4 8 16 31 61 120 236 464 912 1793 3525 6930 ...   |
| 6   | hexanacci    | 1 1 2 4 8 16 32 63 125 248 492 976 1936 3840 7617 ...   |
| 7   | heptanacci   | 1 1 2 4 8 16 32 64 127 253 504 1004 2000 3984 7936 ...  |
| 8   | octonacci    | 1 1 2 4 8 16 32 64 128 255 509 1016 2028 4048 8080 ...  |
| 9   | nonanacci    | 1 1 2 4 8 16 32 64 128 256 511 1021 2040 4076 8144 ...  |
| 10  | decanacci    | 1 1 2 4 8 16 32 64 128 256 512 1023 2045 4088 8172 ...  |

初期値を変えた関連する数列も生成できます：ルーカス数列は $n=2$ のフィボナッチ数列のように直前の2つの値を足しますが、初期値は $\[2, 1]$ です。

# --instructions--

フィボナッチ $n$ ステップ数列とルーカス数列を生成する関数を書いてください。最初のパラメータは $n$、2番目のパラメータは返す要素数、3番目のパラメータはフィボナッチ数列かルーカス数列かを指定します。パラメータが `"f"` の場合はフィボナッチ数列を返し、`"l"` の場合はルーカス数列を返してください。数列は配列として返す必要があります。

# --hints--

`fib_luc` は関数であるべきです。

```js
assert(typeof fib_luc === 'function');
```

`fib_luc(2,10,"f")` は `[1,1,2,3,5,8,13,21,34,55]` を返すべきです。

```js
assert.deepEqual(fib_luc(2, 10, 'f'), globalThis.ans[0]);
```

`fib_luc(3,15,"f")` は `[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136]` を返すべきです。

```js
assert.deepEqual(fib_luc(3, 15, 'f'), globalThis.ans[1]);
```

`fib_luc(4,15,"f")` は `[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536]` を返すべきです。

```js
assert.deepEqual(fib_luc(4, 15, 'f'), globalThis.ans[2]);
```

`fib_luc(2,10,"l")` は `[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76]` を返すべきです。

```js
assert.deepEqual(fib_luc(2, 10, 'l'), globalThis.ans[3]);
```

`fib_luc(3,15,"l")` は `[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ]` を返すべきです。

```js
assert.deepEqual(fib_luc(3, 15, 'l'), globalThis.ans[4]);
```

`fib_luc(4,15,"l")` は `[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ]` を返すべきです。

```js
assert.deepEqual(fib_luc(4, 15, 'l'), globalThis.ans[5]);
```

`fib_luc(5,15,"l")` は `[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]` を返すべきです。

```js
assert.deepEqual(fib_luc(5, 15, 'l'), globalThis.ans[6]);
```

# --before-all--

```js
globalThis.ans = [[1,1,2,3,5,8,13,21,34,55],
[1,1,2,4,7,13,24,44,81,149,274,504,927,1705,3136],
[1,1,2,4,8,15,29,56,108,208,401,773,1490,2872,5536],
[ 2, 1, 3, 4, 7, 11, 18, 29, 47, 76],
[ 2, 1, 3, 6, 10, 19, 35, 64, 118, 217, 399, 734, 1350, 2483, 4567 ],
[ 2, 1, 3, 6, 12, 22, 43, 83, 160, 308, 594, 1145, 2207, 4254, 8200 ],
[ 2, 1, 3, 6, 12, 24, 46, 91, 179, 352, 692, 1360, 2674, 5257, 10335 ]];
```

# --seed--

## --seed-contents--

```js
function fib_luc(n, len, w) {

}
```

# --solutions--

```js
function fib_luc(n, len, w) {
    function nacci(a, n, len) {
        while (a.length < len) {
            let sum = 0;
            for (let i = Math.max(0, a.length - n); i < a.length; i++)
                sum += a[i];
            a.push(sum);
        }
        return a;
    }
    if(w=="f"){
        return nacci(nacci([1,1], n, n), n, len);
    }else{
        return nacci(nacci([2,1], n, n), n, len);
    }
}
```
