---
id: 5992e222d397f00d21122931
title: フィボナッチワード
challengeType: 1
forumTopicId: 302269
dashedName: fibonacci-word
---

# --description--

フィボナッチワードのシーケンスはフィボナッチ数列に類似した方法で作成できますが、連結の繰り返しに焦点を当てています。

<pre>F_Word<sub>1</sub> を <strong>1</strong> と定義します
F_Word<sub>2</sub> を <strong>0</strong> と定義します
F_Word<sub>3</sub> は F_Word<sub>2</sub> と F_Word<sub>1</sub> を連結して作成します。すなわち <strong>01</strong>
F_Word<sub>n</sub> は F_Word<sub>n-1</sub> と F_word<sub>n-2</sub> を連結して作成します
</pre>

このチャレンジではエントロピー計算が必要です。<a href="https://www.freecodecamp.org/learn/coding-interview-prep/rosetta-code/entropy" target="_blank" rel="noopener noreferrer nofollow">こちらのRosetta Codeのチャレンジ</a>を参照してください。

# --instructions--

最初の`n`個のフィボナッチワードを返す関数を書いてください。`n`の数は関数のパラメータとして与えられます。関数はオブジェクトの配列を返す必要があります。オブジェクトは`{ N: 1, Length: 1, Entropy: 0, Word: '1' }`の形式であるべきです。`Entropy`は文字列`Word`に対して計算され、小数点以下8桁に丸められます。このシーケンスのインデックスは`1`から始まることに注意してください。

# --hints--

`fibWord`は関数であるべきです。

```js
assert(typeof fibWord === 'function');
```

`fibWord(5)`は配列を返すべきです。

```js
assert(Array.isArray(fibWord(5)));
```

`fibWord(5)`は`[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }]`を返すべきです。

```js
assert.deepEqual(fibWord(5), globalThis.words5);
```

`fibWord(7)`は`[{ N:1, Length:1, Entropy:0, Word:"1" },{ N:2, Length:1, Entropy:0, Word:"0" },{ N:3, Length:2, Entropy:1, Word:"01" },{ N:4, Length:3, Entropy:0.91829583, Word:"010" },{ N:5, Length:5, Entropy:0.97095059, Word:"01001" }, { N:6, Length:8, Entropy:0.954434, Word:'01001010' }, { N:7, Length:13, Entropy:0.9612366, Word:'0100101001001' }]`を返すべきです。

```js
assert.deepEqual(fibWord(7), globalThis.words7);
```

# --before-all--

```js
globalThis.words5 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' }
];

globalThis.words7 = [
  { N: 1, Length: 1, Entropy: 0, Word: '1' },
  { N: 2, Length: 1, Entropy: 0, Word: '0' },
  { N: 3, Length: 2, Entropy: 1, Word: '01' },
  { N: 4, Length: 3, Entropy: 0.91829583, Word: '010' },
  { N: 5, Length: 5, Entropy: 0.97095059, Word: '01001' },
  { N: 6, Length: 8, Entropy: 0.954434, Word: '01001010' },
  { N: 7, Length: 13, Entropy: 0.9612366, Word: '0100101001001' }
];
```

# --seed--

## --seed-contents--

```js
function fibWord(n) {

}
```

# --solutions--

```js
// Round to digits
function roundFloat(num, digits) {
  return Math.round(num * 10.0**digits) / (10.0**digits);
}

// Entropy calculation for string with only 0 and 1
function entropy(word) {
  function digitEntropy(count) {
    return count < 1 ? 0
      : - count / word.length * Math.log2(count / word.length);
  }
  const numZeros = word.split('').filter(e => e === '0').length;
  const numOnes  = word.length - numZeros;
  return roundFloat(digitEntropy(numZeros) + digitEntropy(numOnes), 8);
}

// Compute array of Fibonacci words
function fibWord(n) {
  return [...Array(n).keys()]
    .reduce((words, i) => {
      const word = i === 0 ? "1"
                 : i === 1 ? "0"
                 : words[i - 1].Word + words[i - 2].Word;
      words.push(
        { N: i + 1, Length: word.length, Entropy: entropy(word), Word: word }
      );
      return words;
    }, []);
}
```
