---
id: 599d1566a02b571412643b84
title: エチオピア乗算
challengeType: 1
forumTopicId: 302257
dashedName: ethiopian-multiplication
---

# --description--

エチオピア乗算は、加算、倍加、半減のみを使って整数を乗算する方法です。

**方法：**

<ol>
  <li>乗算する2つの数を取り、それぞれを2つの列の一番上に書きます</li>
  <li>左の列では、最後の数を繰り返し半分にし、余りは切り捨てて、結果を同じ列の最後の下に書きます。値が<code>1</code>になるまで続けます</li>
  <li>右の列では、最後の数を繰り返し倍にして、その結果を下に書きます。左の列が<code>1</code>の行まで続けます</li>
  <li>作成されたテーブルを調べ、左の列の値が偶数の行は破棄します</li>
  <li>残った右の列の値を合計して、元の2つの数の乗算結果を得ます</li>
</ol>

**例えば：** `17 × 34`

<pre>17   34
</pre>

左の列を半減すると：

<pre>17   34
8
4
2
1
</pre>

右の列を倍加すると：

<pre>17   34
8    68
4   136
2   272
1   544
</pre>

左の列の最初のセルが偶数の行を取り消し線で示すと：

<pre>17   34
8    <strike>68</strike>
4   <strike>136</strike>
2   <strike>272</strike>
1   544
</pre>

右の列の残った数を合計すると：

<!-- markdownlint-disable MD003 -->

<pre>17   34
8    --
4   ---
2   ---
1   544
   ====
    578
</pre>

<!-- markdownlint-enable MD003 -->

つまり、`17` と `34` のエチオピア乗算の結果は `578` です。

# --instructions--

3つの名前付き関数/メソッド/手続き/サブルーチンを定義してください：

<ol>
  <li>整数を半分にする関数、</li>
  <li>整数を倍にする関数、</li>
  <li>整数が偶数かどうかを判定する関数</li>
</ol>

これらの関数を使って、エチオピア乗算を行う関数を作成してください。

<!-- markdownlint-disable MD046-->

# --hints--

`eth_mult` は関数であるべきです。

```js
assert(typeof eth_mult === 'function');
```

`eth_mult(17,34)` は `578` を返すべきです。

```js
assert.equal(eth_mult(17, 34), 578);
```

`eth_mult(23,46)` は `1058` を返すべきです。

```js
assert.equal(eth_mult(23, 46), 1058);
```

`eth_mult(12,27)` は `324` を返すべきです。

```js
assert.equal(eth_mult(12, 27), 324);
```

`eth_mult(56,98)` は `5488` を返すべきです。

```js
assert.equal(eth_mult(56, 98), 5488);
```

`eth_mult(63,74)` は `4662` を返すべきです。

```js
assert.equal(eth_mult(63, 74), 4662);
```

# --seed--

## --seed-contents--

```js
function eth_mult(a, b) {

}
```

# --solutions--

```js
function eth_mult(a, b) {
  let sum = 0; a = [a]; b = [b];

  let half = a => a / 2,
    double = a => a * 2,
    is_even = a => a % 2 == 0;

  while (a[0] !== 1) {
    a.unshift(Math.floor(half(a[0])));
    b.unshift(double(b[0]));
  }

  for (let i = a.length - 1; i > 0; i -= 1) {
    if (!is_even(a[i])) {
      sum += b[i];
    }
  }
  return sum + b[0];
}
```
