---
id: 59c3ec9f15068017c96eb8a3
title: Farey シーケンス
challengeType: 1
forumTopicId: 302266
dashedName: farey-sequence
---

# --description--

Farey シーケンス <code>F<sub>n</sub></code> は、`n` 次の、`0` と `1` の間の既約分数で、分母が `n` 以下のものを小さい順に並べたシーケンスです。

*Farey シーケンス* は時に誤って *Farey シリーズ* と呼ばれることがあります。

各 Farey シーケンスは：

<ul>
  <li>値 0 から始まり、分数で表すと $ \frac{0}{1} $ です</li>
  <li>値 1 で終わり、分数で表すと $ \frac{1}{1}$ です。</li>
</ul>

`1` 次から `5` 次までの Farey シーケンスは以下の通りです：

<ul>
  <li style='list-style: none;'>${\bf\it{F}}_1 = \frac{0}{1}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_2 = \frac{0}{1}, \frac{1}{2}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_3 = \frac{0}{1}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_4 = \frac{0}{1}, \frac{1}{4}, \frac{1}{3}, \frac{1}{2}, \frac{2}{3}, \frac{3}{4}, \frac{1}{1}$</li>
  <li style='list-style: none;'>${\bf\it{F}}_5 = \frac{0}{1}, \frac{1}{5}, \frac{1}{4}, \frac{1}{3}, \frac{2}{5}, \frac{1}{2}, \frac{3}{5}, \frac{2}{3}, \frac{3}{4}, \frac{4}{5}, \frac{1}{1}$</li>
</ul>

# --instructions--

`n` 次の Farey シーケンスを返す関数を書いてください。関数は `n` というパラメータを1つ持ち、シーケンスを配列として返す必要があります。

# --hints--

`farey` は関数であるべきです。

```js
assert(typeof farey === 'function');
```

`farey(3)` は配列を返すべきです。

```js
assert(Array.isArray(farey(3)));
```

`farey(3)` は `['0/1','1/3','1/2','2/3','1/1']` を返すべきです。

```js
assert.deepEqual(farey(3),['0/1', '1/3', '1/2', '2/3', '1/1']);
```

`farey(4)` は `['0/1','1/4','1/3','1/2','2/3','3/4','1/1']` を返すべきです。

```js
assert.deepEqual(farey(4), ['0/1', '1/4', '1/3', '1/2', '2/3', '3/4', '1/1']);
```

`farey(5)` は `['0/1','1/5','1/4','1/3','2/5','1/2','3/5','2/3','3/4','4/5','1/1']` を返すべきです。

```js
assert.deepEqual(farey(5), [
  '0/1',
  '1/5',
  '1/4',
  '1/3',
  '2/5',
  '1/2',
  '3/5',
  '2/3',
  '3/4',
  '4/5',
  '1/1'
]);
```

# --seed--

## --seed-contents--

```js
function farey(n) {

}
```

# --solutions--

```js
function farey(n) {
  const sequence = [{ string: "0/1", float: 0.0 }];
  for (let i = 1; i < n; i++) {
    for (let j = n; j >= i; j--) {
      if (i === 1 || j % i > 0) {
        sequence.push({ string: `${i}/${j}`, float: i / j });
      }
    }
  }
  return sequence
    .sort((a, b) => a.float - b.float)
    .map(e => e.string)
}
```
