---
id: 59e0a8df964e4540d5abe599
title: Brain**** を実行する
challengeType: 1
forumTopicId: 302261
dashedName: execute-brain
---

# --description--

Brain**** インタープリターを実装する関数を書いてください。関数は文字列をパラメータとして受け取り、出力として文字列を返す必要があります。詳細は以下の通りです。

RCBF は、さまざまな言語で Rosetta Code 向けに書かれた<a href="https://rosettacode.org/wiki/Brainf***" target="_blank" rel="noopener noreferrer nofollow">Brain****</a> のコンパイラーおよびインタープリターのセットです。

以下に RCBF の各バージョンへのリンクがあります。

実装では、以下の命令のみ正しく実装すれば十分です。

| コマンド           | 説明                                                                              |
| ----------------- | ---------------------------------------------------------------------------------- |
| <code>></code>    | ポインターを右に移動する                                                          |
| <code>&lt;</code> | ポインターを左に移動する                                                          |
| <code>+</code>    | ポインターが指すメモリセルの値をインクリメントする                                |
| <code>-</code>    | ポインターが指すメモリセルの値をデクリメントする                                  |
| <code>.</code>    | ポインターが指すセルの値に対応する文字を出力する                                  |
| <code>,</code>    | 文字を入力し、ポインターが指すセルに格納する                                      |
| <code>\[</code>   | ポインターが指すセルの値が0の場合、対応する <code>]</code> の次にジャンプする       |
| <code>]</code>    | ポインターが指すセルの値が0以外の場合、対応する <code>\[</code> にジャンプする     |

セルのサイズは任意で、EOF（ファイル終端）サポートはオプションです。メモリが有界か無限かも任意です。

# --hints--

`brain(globalThis.hello)` は文字列を返す必要があります

```js
assert(typeof brain(globalThis.hello) === 'string');
```

`brain("++++++[>++++++++++<-]>+++++.")` は "A" を返す必要があります

```js
assert.equal(brain('++++++[>++++++++++<-]>+++++.'), 'A');
```

`brain(globalThis.hello)` は `Hello World!\n` を返す必要があります

```js
assert.equal(brain(globalThis.hello), 'Hello World!\n');
```

`brain(globalThis.fib)` は `1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89` を返す必要があります

```js
assert.equal(brain(globalThis.fib), '1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89');
```

# --before-all--

```js
globalThis.fib=`+

++

+++

++++

+>+>>

>>++++

+++++++

++++++++

+++++++++

++++++++++

++++++>++++

++++++++++++

+++++++++++++

+++<<<<<<[>[>>

>>>>+>+<<<<<<<-

]>>>>>>>[<<<<<<<

+>>>>>>>-]<[>++++

++++++[-<-[>>+>+<<

<-]>>>[<<<+>>>-]+<[

>[-]<[-]]>[<<[>>>+<<

<-]>>[-]]<<]>>>[>>+>+

<<<-]>>>[<<<+>>>-]+<[>

[-]<[-]]>[<<+>>[-]]<<<<

<<<]>>>>>[++++++++++++++

+++++++++++++++++++++++++

+++++++++.[-]]++++++++++<[

->-<]>+++++++++++++++++++++

+++++++++++++++++++++++++++.

[-]<<<<<<<<<<<<[>>>+>+<<<<-]>

>>>[<<<<+>>>>-]<-[>>.>.<<<[-]]

<<[>>+>+<<<-]>>>[<<<+>>>-]<<[<+

>-]>[<+>-]<<<-]`;
globalThis.hello='++++++++[>++++[>++>+++>+++>+<<<<-]>+>+>->>+[<]<-]>>.>---.+++++++..+++.>>.<-.<.+++.------.--------.>>+.>++.'
```

# --seed--

## --seed-contents--

```js
function brain(prog) {

}
```

# --solutions--

```js
function brain(prog){
  var output="";
    var code; // formatted code
  var ip = 0; // current instruction within code
  var nest = 0; // current bracket nesting (for Out button)
  var ahead = []; // locations of matching brackets

  var data = [0]; // data array (mod by +, -)
  var dp = 0; // index into data (mod by <, >)

  var inp = 0; // current input character (fetch with ,)
  var quit = 0;
    var commands = {
    '>':function() { if (++dp >= data.length) data[dp]=0 },
    '<':function() { if (--dp < 0) quit++ },
    '+':function() { ++data[dp] },
    '-':function() { --data[dp] },
    '[':function() { if (!data[dp]) ip = ahead[ip]; else ++nest },
    ']':function() { if ( data[dp]) ip = ahead[ip]; else --nest },
    ',':function() {
        var c = document.getElementById("input").value.charCodeAt(inp++);
        data[dp] = isNaN(c) ? 0 : c; // EOF: other options are -1 or no change
    },
    '.':function() {
            output+=String.fromCharCode(data[dp]);
            /*var s = document.getElementById("output").innerHTML)
             + String.fromCharCode(data[dp]);
            s = s.replace(/\n/g,"<br>").replace(/ /g,"&amp;nbsp;");
            document.getElementById("output").innerHTML = s;*/
        },
    };

    let ar=prog.split('');
    var st = [], back, error = -1;
    for (ip=0; ip<ar.length; ip++) {
        switch(ar[ip]) {
        case '[':
            st.push(ip);
            break;
        case ']':
            if (st.length == 0) error = ip;
            back = st.pop();
            ahead[ip] = back;
            ahead[back] = ip;
            break;
        }
    }

    for(ip=0;ip<ar.length;ip++){
    if(commands.hasOwnProperty(ar[ip]))
          commands[ar[ip]]();
    }

    return output;
}
```
