---
id: 59f4eafba0343628bb682785
title: Discordianカレンダーの日付
challengeType: 1
forumTopicId: 302250
dashedName: discordian-date
---

# --description--
グレゴリオ暦は、28～31日の12か月からなる太陽暦です。1週間は7日で、1年は52週間と1日で構成されます。1年は365日で、うるう年には2月に1日追加されます。うるう年は4年ごとに訪れますが、400年ごとの最初の3つの世紀年は除かれます。

<a href="http://www.rosettacode.org/wiki/Discordian_date" target="_blank" rel="noopener noreferrer nofollow">Discordianカレンダー</a>はグレゴリオ暦に合わせており、1月1日から始まります。1週間は5日で、1年は73週間あります。1年は365日で、うるう年にはChaos 59とChaos 60の間に1日追加されます。

Discordianカレンダーの季節、曜日、使徒のイベント、祝日は以下の通りです。

季節: `'Chaos', 'Discord', 'Confusion', 'Bureaucracy', 'The Aftermath'`。
  
曜日: `'Sweetmorn', 'Boomtime', 'Pungenday', 'Prickle-Prickle', 'Setting Orange'`。

使徒: `'Mungday', 'Mojoday', 'Syaday', 'Zaraday', 'Maladay'`。

祝日: `'Chaoflux', 'Discoflux', 'Confuflux', 'Bureflux', 'Afflux'`。

# --instructions--

与えられた日付をグレゴリオ暦からDiscordianカレンダーに変換してください。

DiscordianカレンダーのChaos 1、3188 YOLDの日はグレゴリオ暦の2022年1月1日です。

# --hints--

`discordianDate`は関数である必要があります。

```js
assert(typeof discordianDate === 'function');
```

`discordianDate(new Date(2010, 6, 22))`は`"Pungenday, the 57th day of Confusion in the YOLD 3176"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2010, 6, 22)) ===
    'Pungenday, the 57th day of Confusion in the YOLD 3176'
);
```

`discordianDate(new Date(2012, 1, 28))`は`"Prickle-Prickle, the 59th day of Chaos in the YOLD 3178"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2012, 1, 28)) ===
    'Prickle-Prickle, the 59th day of Chaos in the YOLD 3178'
);
```

`discordianDate(new Date(2012, 1, 29))`は`"Setting Orange, the 60th day of Chaos in the YOLD 3178. Celebrate St. Tib\'s Day!"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2012, 1, 29)) ===
    "Setting Orange, the 60th day of Chaos in the YOLD 3178. Celebrate St. Tib's Day!"
);
```

`discordianDate(new Date(2012, 2, 1))`は`"Setting Orange, the 60th day of Chaos in the YOLD 3178"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2012, 2, 1)) ===
    'Setting Orange, the 60th day of Chaos in the YOLD 3178'
);
```

`discordianDate(new Date(2010, 0, 5))`は`"Setting Orange, the 5th day of Chaos in the YOLD 3176. Celebrate Mungday!"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2010, 0, 5)) ===
    'Setting Orange, the 5th day of Chaos in the YOLD 3176. Celebrate Mungday!'
);
```

`discordianDate(new Date(2011, 4, 3))`は`"Pungenday, the 50th day of Discord in the YOLD 3177. Celebrate Discoflux!"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2011, 4, 3)) ===
    'Pungenday, the 50th day of Discord in the YOLD 3177. Celebrate Discoflux!'
);
```

`discordianDate(new Date(2015, 9, 19))`は`"Boomtime, the 73rd day of Bureaucracy in the YOLD 3181"`を返す必要があります。

```js
assert(
  discordianDate(new Date(2015, 9, 19)) ===
    'Boomtime, the 73rd day of Bureaucracy in the YOLD 3181'
);
```

# --seed--

## --seed-contents--

```js
function discordianDate(date) {

  return true;
}
```

# --solutions--

```js
/**
 * All Hail Discordia! - this script prints Discordian date using system date.
 *
 * lang: JavaScript
 * author: jklu
 * contributors: JamesMcGuigan
 *
 * source: https://rosettacode.org/wiki/Discordian_date#JavaScript
 */
const seasons = [
  'Chaos', 'Discord', 'Confusion',
  'Bureaucracy', 'The Aftermath'
];
const weekday = [
  'Sweetmorn', 'Boomtime', 'Pungenday',
  'Prickle-Prickle', 'Setting Orange'
];

const apostle = [
  'Mungday', 'Mojoday', 'Syaday',
  'Zaraday', 'Maladay'
];

const holiday = [
  'Chaoflux', 'Discoflux', 'Confuflux',
  'Bureflux', 'Afflux'
];


Date.prototype.isLeapYear = function() {
  const year = this.getFullYear();
  if ((year & 3) !== 0) { return false; }
  return ((year % 100) !== 0 || (year % 400) === 0);
};

// Get Day of Year
Date.prototype.getDOY = function() {
  const dayCount = [0, 31, 59, 90, 120, 151, 181, 212, 243, 273, 304, 334];
  const mn = this.getMonth();
  const dn = this.getDate();
  let dayOfYear = dayCount[mn] + dn;
  if (mn > 1 && this.isLeapYear()) { dayOfYear += 1; }
  return dayOfYear;
};

Date.prototype.isToday = function() {
  const today = new Date();
  return this.getDate() === today.getDate()
      && this.getMonth() === today.getMonth()
      && this.getFullYear() === today.getFullYear()
  ;
};

function discordianDate(date) {
  if (!date) { date = new Date(); }

  const y = date.getFullYear();
  const yold = y + 1166;
  let dayOfYear = date.getDOY();
  let celebrateHoliday = null;

  if (date.isLeapYear()) {
    if (dayOfYear === 60) {
      celebrateHoliday = 'St. Tib\'s Day';
    }
    else if (dayOfYear > 60) {
      dayOfYear--;
    }
  }
  dayOfYear--;

  const divDay = Math.floor(dayOfYear / 73);

  const seasonDay = (dayOfYear % 73) + 1;
  if (seasonDay === 5) {
    celebrateHoliday = apostle[divDay];
  }
  if (seasonDay === 50) {
    celebrateHoliday = holiday[divDay];
  }

  const season = seasons[divDay];
  const dayOfWeek = weekday[dayOfYear % 5];

  const nth = (seasonDay % 10 === 1) ? 'st'
          : (seasonDay % 10 === 2) ? 'nd'
          : (seasonDay % 10 === 3) ? 'rd'
                                  : 'th';

  return ''
         + dayOfWeek
         + ', the ' + seasonDay + nth
         + ' day of ' + season
         + ' in the YOLD ' + yold
         + (celebrateHoliday ? '. Celebrate ' + celebrateHoliday + '!' : '')
    ;
}
```
