---
id: 5a23c84252665b21eecc7eb1
title: 単位行列
challengeType: 1
forumTopicId: 302290
dashedName: identity-matrix
---

# --description--

*単位行列*とは、サイズが \\( n \times n \\) の正方行列で、対角要素がすべて `1`（1）で、その他の要素がすべて `0`（0）である行列です。

<ul>
  <li style='list-style: none;'>\(\displaystyle I_{n}=\begin{bmatrix} 1 &#x26; 0 &#x26; 0 \cr 0 &#x26; 1 &#x26; 0 \cr 0 &#x26; 0 &#x26; 1 \cr \end{bmatrix}\)</li>
</ul>

# --instructions--

数値 `n` をパラメータとして受け取り、\\( n \times n \\) の単位行列を返す関数を書いてください。

# --before-each--

```js
let results=[[ [ 1 ] ],
[ [ 1, 0 ], [ 0, 1 ] ],
[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ],
[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]]
```

# --hints--

`idMatrix` は関数であるべきです。

```js
assert(typeof idMatrix == 'function');
```

`idMatrix(1)` は配列を返すべきです。

```js
assert(Array.isArray(idMatrix(1)));
```

`idMatrix(1)` は `[ [ 1 ] ]` を返すべきです。

```js
assert.deepEqual(idMatrix(1), results[0]);
```

`idMatrix(2)` は `[ [ 1, 0 ], [ 0, 1 ] ]` を返すべきです。

```js
assert.deepEqual(idMatrix(2), results[1]);
```

`idMatrix(3)` は `[ [ 1, 0, 0 ], [ 0, 1, 0 ], [ 0, 0, 1 ] ]` を返すべきです。

```js
assert.deepEqual(idMatrix(3), results[2]);
```

`idMatrix(4)` は `[ [ 1, 0, 0, 0 ], [ 0, 1, 0, 0 ], [ 0, 0, 1, 0 ], [ 0, 0, 0, 1 ] ]` を返すべきです。

```js
assert.deepEqual(idMatrix(4), results[3]);
```

# --seed--

## --seed-contents--

```js
function idMatrix(n) {

}
```

# --solutions--

```js
function idMatrix(n) {
    return Array.apply(null, new Array(n)).map(function (x, i, xs) {
        return xs.map(function (_, k) {
            return i === k ? 1 : 0;
        })
    });
}
```
