---
id: 5a23c84252665b21eecc7ee0
title: 左階乗
challengeType: 1
forumTopicId: 302302
dashedName: left-factorials
---

# --description--

**左階乗**、$ !n $ は、*部分階乗* または *階乗和* のいずれかを指す場合があります。同じ記法が2つの異なる定義で混同して使われることがあります。時には、*部分階乗*（*置換なし順列*とも呼ばれます）は以下のいずれかの記法を使うことがあります：

<ul>
  <li>$!n`$</li>
  <li>$!n$</li>
  <li>$n¡$</li>
</ul>

（見た目にはわかりにくいですが、最後の例は逆さまの感嘆符を使っています。）この課題では、**左階乗**の計算に次の式を使います：

$ !n = \\sum\_{k=0}^{n-1} k! $

ここで $!0 = 0$ です。

# --instructions--

与えられた数の左階乗を計算する関数を書いてください。

# --hints--

`leftFactorial` は関数である必要があります。

```js
assert(typeof leftFactorial == 'function');
```

`leftFactorial(0)` は数値を返す必要があります。

```js
assert(typeof leftFactorial(0) == 'number');
```

`leftFactorial(0)` は `0` を返す必要があります。

```js
assert.equal(leftFactorial(0), 0);
```

`leftFactorial(1)` は `1` を返す必要があります。

```js
assert.equal(leftFactorial(1), 1);
```

`leftFactorial(2)` は `2` を返す必要があります。

```js
assert.equal(leftFactorial(2), 2);
```

`leftFactorial(3)` は `4` を返す必要があります。

```js
assert.equal(leftFactorial(3), 4);
```

`leftFactorial(10)` は `409114` を返す必要があります。

```js
assert.equal(leftFactorial(10), 409114);
```

`leftFactorial(17)` は `22324392524314` を返す必要があります。

```js
assert.equal(leftFactorial(17), 22324392524314);
```

`leftFactorial(19)` は `6780385526348314` を返す必要があります。

```js
assert.equal(leftFactorial(19), 6780385526348314);
```

# --seed--

## --seed-contents--

```js
function leftFactorial(n) {

}
```

# --solutions--

```js
function leftFactorial(n) {
  if (n == 0) return 0;
  if (n == 1) return 1;

  // Note: for n>=20, the result may not be correct.
  // This is because JavaScript uses 53 bit integers and
  // for n>=20 result becomes too large.

  let res = 2,
    fact = 2;
  for (var i = 2; i < n; i++) {
    res += fact;
    fact *= i + 1;
  }

  return res;
}
```
