---
id: 5a23c84252665b21eecc8010
title: ソートアルゴリズム/Shellソート
challengeType: 1
forumTopicId: 302317
dashedName: sorting-algorithmsshell-sort
---

# --description--

Shellソートアルゴリズムを使って要素の配列をソートする関数を書いてください。関数はソート済みの配列を返す必要があります。

Shellソート（ShellsortまたはShellの方法とも呼ばれます）は、発明者のDonald Shellにちなんで名付けられ、1959年にアルゴリズムが公開されました。

Shellソートは、増分シーケンスに基づく複数の挿入ソートを組み合わせたものです。各パスの後に増分サイズが減少し、最終的に増分サイズが1になります。

増分サイズが1のときは基本的な挿入ソートですが、その時点でデータはほぼソート済みであることが保証されているため、挿入ソートの「最良ケース」となります。

増分シーケンスは1で終わればどんなものでもデータをソートできますが、実際にはより効果的なものがあります。

経験的な研究では、約2.2の比率の幾何学的増分シーケンスが実用的に良い結果を出すことが示されています。

# --hints--

`shellSort` は関数である必要があります。

```js
assert(typeof shellSort == 'function');
```

`shellSort([25, 32, 12, 7, 20])` は配列を返す必要があります。

```js
assert(Array.isArray(shellSort([25, 32, 12, 7, 20])));
```

`shellSort([25, 32, 12, 7, 20])` は `[7, 12, 20, 25, 32]` を返す必要があります。

```js
assert.deepEqual(shellSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`shellSort([38, 45, 35, 8, 13])` は `[8, 13, 35, 38, 45]` を返す必要があります。

```js
assert.deepEqual(shellSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`shellSort([43, 36, 20, 34, 24])` は `[20, 24, 34, 36, 43]` を返す必要があります。

```js
assert.deepEqual(shellSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`shellSort([12, 33, 26, 18, 1, 16, 38])` は `[1, 12, 16, 18, 26, 33, 38]` を返す必要があります。

```js
assert.deepEqual(shellSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`shellSort([3, 39, 48, 16, 1, 4, 29])` は `[1, 3, 4, 16, 29, 39, 48]` を返す必要があります。

```js
assert.deepEqual(shellSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function shellSort(a) {

}
```

# --solutions--

```js
function shellSort(a) {
  for (var h = a.length; h > 0; h = parseInt(h / 2)) {
    for (var i = h; i < a.length; i++) {
      var k = a[i];
      for (var j = i; j >= h && k < a[j - h]; j -= h) a[j] = a[j - h];
      a[j] = k;
    }
  }
  return a;
}
```
