---
id: 5a23c84252665b21eecc7ec2
title: Jaro距離
challengeType: 1
forumTopicId: 302292
dashedName: jaro-distance
---

# --description--

Jaro距離は、2つの文字列間の類似度を測る指標です。Jaro距離が高いほど、文字列はより類似しています。スコアは正規化されており、`0`が類似度ゼロ、`1`が完全一致を表します。

**定義**

2つの文字列 \\(s_1\\) と \\(s_2\\) のJaro距離 \\( d_j \\) は次のように定義されます。

\\begin{align}d_j = \\begin{cases}0& & \\text{if }m=0 \\\\\\\\{\\frac {1}{3}}\\left({\\frac {m}{|s\_{1}|}}+{\\frac {m}{|s\_{2}|}}+{\\frac {m-t}{m}}\\right)& & \\text{otherwise}\\end{cases}\\end{align}

ここで：

<ul>
  <li>\(m\) は<i>一致する文字</i>の数です；</li>
  <li>\(t\) は<i>転置</i>の半数です。</li>
</ul>

\\(s_1\\) と \\(s_2\\) のそれぞれの文字は、同じであり、かつ \\(\\left\\lfloor\\frac{\\max(|s_1|,|s_2|)}{2}\\right\\rfloor-1\\) 以内の距離にある場合にのみ*一致*とみなされます。

\\(s_1\\) の各文字は、\\(s_2\\) の一致するすべての文字と比較されます。順序が異なる一致文字の数を2で割ったものが*転置*の数となります。

**例**

文字列 \\(s_1\\) *DWAYNE* と \\(s_2\\) *DUANE* の場合：

<ul>
  <li>\(m = 4\)</li>
  <li>\(|s_1| = 6\)</li>
  <li>\(|s_2| = 5\)</li>
  <li>\(t = 0\)</li>
</ul>

Jaroスコアは次のように計算されます：\\(d_j = \\frac{1}{3}\\left(\\frac{4}{6} + \\frac{4}{5} + \\frac{4-0}{4}\\right) = 0.822\\)。

# --instructions--

2つの文字列をパラメータとして受け取り、対応するJaro距離を返す関数aを作成してください。

# --hints--

`jaro`は関数である必要があります。

```js
assert(typeof jaro == 'function');
```

`jaro("MARTHA", "MARHTA")`は数値を返す必要があります。

```js
assert(typeof jaro('MARTHA', 'MARHTA') == 'number');
```

`jaro("MARTHA", "MARHTA")`は`0.9444444444444445`を返す必要があります。

```js
assert.equal(jaro('MARTHA', 'MARHTA'), 0.9444444444444445);
```

`jaro("DIXON", "DICKSONX")`は`0.7666666666666666`を返す必要があります。

```js
assert.equal(jaro('DIXON', 'DICKSONX'), 0.7666666666666666);
```

`jaro("JELLYFISH", "SMELLYFISH")`は`0.8962962962962964`を返す必要があります。

```js
assert.equal(jaro('JELLYFISH', 'SMELLYFISH'), 0.8962962962962964);
```

`jaro("HELLOS", "CHELLO")`は`0.888888888888889`を返す必要があります。

```js
assert.equal(jaro('HELLOS', 'CHELLO'), 0.888888888888889);
```

`jaro("ABCD", "BCDA")`は`0.8333333333333334`を返す必要があります。

```js
assert.equal(jaro('ABCD', 'BCDA'), 0.8333333333333334);
```

# --seed--

## --seed-contents--

```js
function jaro(s, t) {

}
```

# --solutions--

```js
function jaro(s, t) {
  var s_len = s.length;
  var t_len = t.length;

  if (s_len == 0 && t_len == 0) return 1;

  var match_distance = Math.max(s_len, t_len) / 2 - 1;

  var s_matches = new Array(s_len);
  var t_matches = new Array(t_len);

  var matches = 0;
  var transpositions = 0;

  for (var i = 0; i < s_len; i++) {
    var start = Math.max(0, i - match_distance);
    var end = Math.min(i + match_distance + 1, t_len);

    for (var j = start; j < end; j++) {
      if (t_matches[j]) continue;
      if (s.charAt(i) != t.charAt(j)) continue;
      s_matches[i] = true;
      t_matches[j] = true;
      matches++;
      break;
    }
  }

  if (matches == 0) return 0;

  var k = 0;
  for (var i = 0; i < s_len; i++) {
    if (!s_matches[i]) continue;
    while (!t_matches[k]) k++;
    if (s.charAt(i) != t.charAt(k)) transpositions++;
    k++;
  }

  return ((matches / s_len) +
    (matches / t_len) +
    ((matches - transpositions / 2.0) / matches)) / 3.0;
}
```
