---
id: 5a23c84252665b21eecc8028
title: スターン・ブロコットシーケンス
challengeType: 1
forumTopicId: 302324
dashedName: stern-brocot-sequence
---

# --description--

この課題では、スターン・ブロコットシーケンスを、<a href="https://rosettacode.org/wiki/Fibonacci_sequence" target="_blank" rel="noopener noreferrer nofollow">フィボナッチシーケンス</a>の生成に使われるアルゴリズムに似た方法で生成します。

<ol>
  <li>シーケンスの最初と2番目の要素はどちらも1です：</li>
    <ul><li>1, 1</li></ul>
  <li>シーケンスの2番目の要素を考えます</li>
  <li>考えている要素とその前の要素の和を計算し、(1 + 1) = 2、それをシーケンスの末尾に追加します：</li>
    <ul><li>1, 1, 2</li></ul>
  <li>考えている要素をシーケンスの末尾に追加します：</li>
    <ul><li>1, 1, 2, 1</li></ul>
  <li>次の要素（3番目の要素、つまり2）を考えます</li>
  <li>3に戻ります</li>
    <ul>
      <li></li>
      <li> ─── もう一回ループを展開すると： ───</li>
      <li></li>
    </ul>
  <li>考えている要素とその前の要素の和を計算し、(2 + 1) = 3、それをシーケンスの末尾に追加します：</li>
    <ul><li>1, 1, 2, 1, 3</li></ul>
  <li>考えている要素をシーケンスの末尾に追加します：</li>
    <ul><li>1, 1, 2, 1, 3, 2</li></ul>
  <li>次の要素（4番目の要素、つまり1）を考えます</li>
</ol>

# --instructions--

スターン・ブロコットシーケンスの中で、$ n $ が初めて現れる位置を返す関数を作成してください。このシーケンスは1から始まるインデックスを使っています。

# --hints--

`sternBrocot` は関数であるべきです。

```js
assert(typeof sternBrocot == 'function');
```

`sternBrocot(2)` は数値を返すべきです。

```js
assert(typeof sternBrocot(2) == 'number');
```

`sternBrocot(2)` は `3` を返すべきです。

```js
assert.equal(sternBrocot(2), 3);
```

`sternBrocot(3)` は `5` を返すべきです。

```js
assert.equal(sternBrocot(3), 5);
```

`sternBrocot(5)` は `11` を返すべきです。

```js
assert.equal(sternBrocot(5), 11);
```

`sternBrocot(7)` は `19` を返すべきです。

```js
assert.equal(sternBrocot(7), 19);
```

`sternBrocot(10)` は `39` を返すべきです。

```js
assert.equal(sternBrocot(10), 39);
```

# --seed--

## --seed-contents--

```js
function sternBrocot(num) {

}
```

# --solutions--

```js
function sternBrocot(num) {
  function f(n) {
    return n < 2
      ? n
      : n & 1
      ? f(Math.floor(n / 2)) + f(Math.floor(n / 2 + 1))
      : f(Math.floor(n / 2));
  }

  function gcd(a, b) {
    return a ? (a < b ? gcd(b % a, a) : gcd(a % b, b)) : b;
  }
  var n;
  for (n = 1; f(n) != num; n++);
  return n;
}
```
