---
id: 5a23c84252665b21eecc7e7b
title: ジェネレーターと指数関数
challengeType: 1
forumTopicId: 302275
dashedName: generatorexponential
---

# --description--

ジェネレーターとは、関数やプロシージャのような実行可能なエンティティで、値のシーケンスを1つずつ生成するコードを含んでいます。ジェネレーターを呼び出すたびに、シーケンスの次の値が提供されます。

ジェネレーターは、内部状態が「自然に」管理されるように、コルーチンやオブジェクトの上に構築されることが多いです。

ジェネレーターは、シーケンスが潜在的に無限であり、シーケンスの次の値を最小限の状態で構築できる場合によく使われます。

# --instructions--

ジェネレーターを使って平方数と立方数を生成する関数を書いてください。平方数のジェネレーターからすべての立方数をフィルターする新しいジェネレーターを作成してください。

関数は、フィルターされたジェネレーターの\\( n^{th} \\)の値を返す必要があります。

例えば\\(n=7\\)の場合、シーケンスは4、9、16、25、36、49、81となり、関数は81を返します。ここで64は立方数なのでフィルターされます。

# --hints--

`exponentialGenerator`は関数である必要があります。

```js
assert(typeof exponentialGenerator == 'function');
```

`exponentialGenerator()`は数値を返す必要があります。

```js
assert(typeof exponentialGenerator(10) == 'number');
```

`exponentialGenerator(10)`は`144`を返す必要があります。

```js
assert.equal(exponentialGenerator(10), 144);
```

`exponentialGenerator(12)`は`196`を返す必要があります。

```js
assert.equal(exponentialGenerator(12), 196);
```

`exponentialGenerator(14)`は`256`を返す必要があります。

```js
assert.equal(exponentialGenerator(14), 256);
```

`exponentialGenerator(20)`は`484`を返す必要があります。

```js
assert.equal(exponentialGenerator(20), 484);
```

`exponentialGenerator(25)`は`784`を返す必要があります。

```js
assert.equal(exponentialGenerator(25), 784);
```

# --seed--

## --seed-contents--

```js
function exponentialGenerator(n) {

}
```

# --solutions--

```js
function exponentialGenerator(n){
  function* PowersGenerator(m) {
    var n=0;
    while(1) {
        yield Math.pow(n, m);
        n += 1;
    }
  }

  function* FilteredGenerator(g, f){
    var value = g.next().value;
    var filter = f.next().value;
    while(1) {
        if( value < filter ) {
            yield value;
            value = g.next().value;
        } else if ( value > filter ) {
            filter = f.next().value;
        } else {
            value = g.next().value;
            filter = f.next().value;
        }
    }
  }

  var squares = PowersGenerator(2);
  var cubes = PowersGenerator(3);

  var filtered = FilteredGenerator(squares, cubes);

  var curr=0;
  for(var i=0;i<n;i++) curr=filtered.next();

  return curr.value;
}
```
