---
id: 5a23c84252665b21eecc7eca
title: カプレカ数
challengeType: 1
forumTopicId: 302296
dashedName: kaprekar-numbers
---

# --description--

正の整数がカプレカ数である条件は以下の通りです：

<ul>
  <li>その数が1である、または、</li>
  <li>その数の平方の10進表記を一度だけ2つの正の整数の部分に分割でき、その2つの部分の和が元の数になること。</li>
</ul>

ただし、部分の一方がすべて0で構成される分割は無効です。0は正の整数とはみなされません。

カプレカ数の例：

<ul>
  <li><code>2223</code> はカプレカ数です。なぜなら <code>2223 * 2223 = 4941729</code> であり、<code>4941729</code> を <code>494</code> と <code>1729</code> に分割でき、<code>494 + 1729 = 2223</code> となるからです。</li>
  <li>カプレカ数の列は A006886 として知られており、<code>1, 9, 45, 55, ...</code> から始まります。</li>
</ul>

# --instructions--

数値 $n$ と基数 $bs$ を受け取り、その基数において $n$ がカプレカ数であれば true を返し、そうでなければ false を返す関数を書いてください。

# --hints--

`isKaprekar` は関数であるべきです。

```js
assert(typeof isKaprekar == 'function');
```

`isKaprekar(1, 10)` はブール値を返すべきです。

```js
assert(typeof isKaprekar(1, 10) == 'boolean');
```

`isKaprekar(1, 10)` は `true` を返すべきです。

```js
assert.equal(isKaprekar(1, 10), true);
```

`isKaprekar(9, 10)` は `true` を返すべきです。

```js
assert.equal(isKaprekar(9, 10), true);
```

`isKaprekar(2223, 10)` は `true` を返すべきです。

```js
assert.equal(isKaprekar(2223, 10), true);
```

`isKaprekar(22823, 10)` は `false` を返すべきです。

```js
assert.equal(isKaprekar(22823, 10), false);
```

`isKaprekar(9, 17)` は `false` を返すべきです。

```js
assert.equal(isKaprekar(9, 17), false);
```

`isKaprekar(225, 17)` は `true` を返すべきです。

```js
assert.equal(isKaprekar(225, 17), true);
```

`isKaprekar(999, 17)` は `false` を返すべきです。

```js
assert.equal(isKaprekar(999, 17), false);
```

# --seed--

## --seed-contents--

```js
function isKaprekar(n, bs) {

}
```

# --solutions--

```js
function isKaprekar(n, bs) {
  if (n < 1) return false;
  if (n == 1) return true;
  for (var a = n * n, b = 0, s = 1; a; s *= bs) {
    b += (a % bs) * s;
    a = Math.floor(a / bs);
    if (b && a + b == n) return true;
  }
  return false;
}
```
