---
id: 5a23c84252665b21eecc8041
title: シリーズの和
challengeType: 1
forumTopicId: 302333
dashedName: sum-of-a-series
---

# --description--

<em>シリーズ</em>の**n**<sup>番目</sup>の項、つまり対応する<em>シーケンス</em>の最初の**n**項の和を計算してください。非公式には、この値や**n**が無限大に近づくときの極限は*シリーズの和*とも呼ばれ、この課題のタイトルになっています。この課題では、$S_n = \displaystyle\sum_{k=1}^n \frac{1}{k^2}$ を使います。

# --instructions--

$a$と$b$をパラメータとして受け取り、シーケンスの$a$番目から$b$番目までの項の和を返す関数を書いてください。

# --hints--

`sum` は関数である必要があります。

```js
assert(typeof sum == 'function');
```

`sum(1, 100)` は数値を返す必要があります。

```js
assert(typeof sum(1, 100) == 'number');
```

`sum(1, 100)` は `1.6349839001848923` を返す必要があります。

```js
assert.equal(sum(1, 100), 1.6349839001848923);
```

`sum(33, 46)` は `0.009262256361481223` を返す必要があります。

```js
assert.equal(sum(33, 46), 0.009262256361481223);
```

`sum(21, 213)` は `0.044086990748706555` を返す必要があります。

```js
assert.equal(sum(21, 213), 0.044086990748706555);
```

`sum(11, 111)` は `0.08619778593108679` を返す必要があります。

```js
assert.equal(sum(11, 111), 0.08619778593108679);
```

`sum(1, 10)` は `1.5497677311665408` を返す必要があります。

```js
assert.equal(sum(1, 10), 1.5497677311665408);
```

# --seed--

## --seed-contents--

```js
function sum(a, b) {

}
```

# --solutions--

```js
function sum(a, b) {
  function fn(x) {
    return 1 / (x * x);
  }
  var s = 0;
  for (; a <= b; a++) s += fn(a);
  return s;
}
```
