---
id: 5a23c84252665b21eecc8045
title: サザーランド・ホジマン多角形クリッピング
challengeType: 1
forumTopicId: 302336
dashedName: sutherland-hodgman-polygon-clipping
---

# --description--

サザーランド・ホジマンのクリッピングアルゴリズムは、任意の多角形（「対象多角形」）と凸多角形（「クリップ多角形」）の交差部分となる多角形を見つけます。これはコンピューターグラフィックス（特に2Dグラフィックス）で、表示する必要のない多角形の部分を除去してシーンの複雑さを減らすために使われます。次の点で定義される閉じた多角形を考えてください：

<pre>[(50, 150), (200, 50), (350, 150), (350, 300), (250, 300), (200, 250), (150, 350), (100, 250), (100, 200)]</pre>

そして、次の点で定義される長方形でクリップしてください：

<pre>[(100, 100), (300, 100), (300, 300), (100, 300)]</pre>

# --instructions--

2つの配列をパラメータとして受け取る関数を書いてください。最初の配列は対象多角形の点を含み、2番目の配列はクリップ多角形の点を含みます。関数はクリップされた多角形の点を含む配列を返す必要があります。各数値は小数点以下3桁に丸めてください。

# --before-each--

```js
const subjectPolygon1 = [
  [50, 150],
  [200, 50],
  [350, 150],
  [350, 300],
  [250, 300],
  [200, 250],
  [150, 350],
  [100, 250],
  [100, 200]
];
const clipPolygon1 = [
  [100, 100],
  [300, 100],
  [300, 300],
  [100, 300]
];
const expected1 = [
  [100, 116.667],
  [125, 100],
  [275, 100],
  [300, 116.667],
  [300, 300],
  [250, 300],
  [200, 250],
  [175, 300],
  [125, 300],
  [100, 250]
];

const subjectPolygon2 = [
  [150, 200],
  [400, 450],
  [30, 50]
];
const clipPolygon2 = [
  [10, 10],
  [300, 200],
  [400, 600],
  [100, 300]
];
const expected2 = [
  [150, 200],
  [350, 400],
  [348.611, 394.444],
  [30, 50]
];

const subjectPolygon3 = [
  [250, 200],
  [100, 450],
  [130, 250]
];
const clipPolygon3 = [
  [50, 60],
  [100, 230],
  [400, 600],
  [100, 300]
];
const expected3 = [
  [129.167, 329.167],
  [119.565, 319.565],
  [121.854, 304.305]
];
```

# --hints--

`clip` は関数であるべきです。

```js
assert(typeof clip == 'function');
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` は配列を返すべきです。

```js
assert(Array.isArray(clip(subjectPolygon1, clipPolygon1)));
```

`clip([[50, 150], [200, 50], [350, 150], [350, 300], [250, 300], [200, 250], [150, 350], [100, 250], [100, 200]], [[100, 100], [300, 100], [300, 300], [100, 300]])` は `[[100, 116.667], [125, 100], [275, 100], [300, 116.667], [300, 300], [250, 300], [200, 250], [175, 300], [125, 300], [100, 250]]` を返すべきです。

```js
assert.deepEqual(clip(subjectPolygon1, clipPolygon1), expected1);
```

`clip([[150, 200], [400, 450], [30, 50]], [[10, 10], [300, 200], [400, 600], [100, 300]])` は `[[150, 200], [350, 400], [348.611, 394.444], [30, 50]]` を返すべきです。

```js
assert.deepEqual(clip(subjectPolygon2, clipPolygon2), expected2);
```

`clip([[250, 200], [100, 450], [130, 250]], [[50, 60], [100, 230], [400, 600], [100, 300]])` は `[[129.167, 329.167], [119.565, 319.565], [121.854, 304.305]]` を返すべきです。

```js
assert.deepEqual(clip(subjectPolygon3, clipPolygon3), expected3);
```

# --seed--

## --seed-contents--

```js
function clip(subjectPolygon, clipPolygon) {

}
```

# --solutions--

```js
function clip(subjectPolygon, clipPolygon) {
  var cp1, cp2, s, e, i, j;
  var inside = function(p) {
    return (
      (cp2[0] - cp1[0]) * (p[1] - cp1[1]) > (cp2[1] - cp1[1]) * (p[0] - cp1[0])
    );
  };
  var intersection = function() {
    var dc = [cp1[0] - cp2[0], cp1[1] - cp2[1]],
      dp = [s[0] - e[0], s[1] - e[1]],
      n1 = cp1[0] * cp2[1] - cp1[1] * cp2[0],
      n2 = s[0] * e[1] - s[1] * e[0],
      n3 = 1.0 / (dc[0] * dp[1] - dc[1] * dp[0]);
    return [(n1 * dp[0] - n2 * dc[0]) * n3, (n1 * dp[1] - n2 * dc[1]) * n3];
  };
  var outputList = subjectPolygon;
  cp1 = clipPolygon[clipPolygon.length - 1];
  for (j in clipPolygon) {
    var cp2 = clipPolygon[j];
    var inputList = outputList;
    outputList = [];
    s = inputList[inputList.length - 1]; //last on the input list
    for (i in inputList) {
      var e = inputList[i];
      if (inside(e)) {
        if (!inside(s)) {
          outputList.push(intersection());
        }
        outputList.push(e);
      } else if (inside(s)) {
        outputList.push(intersection());
      }
      s = e;
    }
    cp1 = cp2;
  }
  return outputList.map(e => e.map(f => Math.round(f * 1000) / 1000));
}
```
