---
id: 5a23c84252665b21eecc8038
title: Subleq
challengeType: 1
forumTopicId: 302328
dashedName: subleq
---

# --description--

Subleqはワンインストラクションセットコンピューター（OISC）の一例です。

名前は唯一の命令である、**SU**btract and **B**ranch if **L**ess than or **EQ**ual to zero（引き算して、結果が0以下なら分岐）に由来します。

あなたの課題は、このようなマシンをエミュレートするインタープリターを作成することです。

マシンのメモリは符号付き整数の配列で構成されます。適切なワードサイズであれば問題ありませんが、メモリは負の数と正の数の両方を保持できる必要があります。

実行は命令ポインターが最初のワード、つまりアドレス0を指すところから始まります。実行の流れは以下の通りです：

<ol>
  <li>A、B、Cを命令ポインターから連続する3つのメモリワードに格納されている値とします。</li>
  <li>命令ポインターを3ワード進めて、Cを含むアドレスの次のアドレスを指すようにします。</li>
  <li>Aが-1の場合、標準入力から1文字読み込み、そのコードポイントをBで指定されたアドレスに格納します。Cは使いません。</li>
  <li>Bが-1の場合、Aで指定されたアドレスにある数値をコードポイントとして解釈し、対応する文字を出力します。Cは使いません。</li>
  <li>それ以外の場合、AとBはメモリのアドレスとして扱われます。Aで指定されたアドレスの数値をBで指定されたアドレスの数値から引き、その結果をBのアドレスに格納します。結果が0以下なら、Cの値が新しい命令ポインターになります。</li>
  <li>命令ポインターが負の値になったら、実行を停止します。</li>
</ol>

-1以外の負のアドレスは-1と同等に扱っても、エラーにしても構いません。

あなたの解答は、マシン上で実行するプログラムと、そのプログラムに与える入力を別々に受け取るようにしてください。

このプログラムは生のsubleq「マシンコード」で、空白区切りの10進数の数値のみで構成され、シンボリック名やアセンブリレベルの拡張はありません。アドレス0からメモリにロードされます。この「Hello, world!」プログラムを入力したときの出力を示してください。（例はASCIIまたはそのスーパーセット（Latin-N文字セットやUnicodeなど）を想定しています。非ASCII互換環境の場合は別の文字セットに翻訳しても構いません。）

<pre>15 17 -1 17 -1 -1 16 1 -1 16 3 -1 15 15 0 0 -1 72 101 108 108 111 44 32 119 111 114 108 100 33 10 0</pre>

これは仮想的なアセンブラ言語で以下のようなものに対応します：

<pre>start:
    zero, message, -1
    message, -1, -1
    neg1, start+1, -1
    neg1, start+3, -1
    zero, zero, start
zero: 0
neg1: -1
message: "Hello, world!\n\0"
</pre>

# --instructions--

整数の配列をパラメータとして受け取る関数を書いてください。これはメモリ要素を表します。関数はこのシーケンスを解釈し、出力文字列を返す必要があります。この課題では標準入力はないものとします。

# --hints--

`Subleq` は関数であるべきです。

```js
assert(typeof Subleq == 'function');
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])` は文字列を返すべきです。

```js
assert(
  typeof Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]) == 'string'
);
```

`Subleq([15, 17, -1, 17, -1, -1, 16, 1, -1, 16, 3, -1, 15, 15, 0, 0, -1, 72, 101, 108, 108, 111, 44, 32, 119, 111, 114, 108, 100, 33, 0])` は `"Hello, world!"` を返すべきです。

```js
assert.equal(
  Subleq([
    15,
    17,
    -1,
    17,
    -1,
    -1,
    16,
    1,
    -1,
    16,
    3,
    -1,
    15,
    15,
    0,
    0,
    -1,
    72,
    101,
    108,
    108,
    111,
    44,
    32,
    119,
    111,
    114,
    108,
    100,
    33,
    0
  ]),
  'Hello, world!'
);
```

# --seed--

## --seed-contents--

```js
function Subleq(mem) {

}
```

# --solutions--

```js
function Subleq(mem) {
  var out = '';
  var instructionPointer = 0;
  do {
    var a = mem[instructionPointer];
    var b = mem[instructionPointer + 1];
    if (a === -1) {
    } else if (b === -1) {
      out += String.fromCharCode(mem[a]);
    } else {
      mem[b] -= mem[a];
      if (mem[b] < 1) {
        instructionPointer = mem[instructionPointer + 2];
        continue;
      }
    }
    instructionPointer += 3;
  } while (instructionPointer >= 0);

  return out;
}
```
