---
id: 5a23c84252665b21eecc8013
title: ソートアルゴリズム/ストランドソート
challengeType: 1
forumTopicId: 302319
dashedName: sorting-algorithmsstrand-sort
---

# --description--

**ストランドソート**は、ソート済みの部分集合を作成し、それらをマージして最終結果を作ります。

`unsortedArray = [3, 1, 4, 2]`を考えてみましょう。最初のアイテム`3`を選び、別の配列にコピーします。このアイテムの後にあるより大きなアイテムを探します。ここでは`4`が見つかった場合、それを別の配列`[3, 4]`にコピーし、次のアイテムをこの新しい値`4`と比較します。

配列の最後まで到達したら、コピーしたアイテム`[3, 4]`を削除し、`unsortedArray`に残っている最初のアイテム、ここでは`1`から再開します。

このプロセスを繰り返すと、2つのソート済み配列`[3, 4]`と`[1, 2]`が得られます。これら2つの配列をマージして`strandSortedArray`を作成します。

```js
const unsortedArray = [3, 1, 4, 2];
const strandsortedArray = [1, 2, 3, 4];
```

**ストランドソート**を使って配列をソートする関数を書いてください。関数はソート済みの配列を返す必要があります。


# --hints--

`strandSort`は関数であるべきです。

```js
assert(typeof strandSort == 'function');
```

`strandSort([25, 32, 12, 7, 20])`は配列を返すべきです。

```js
assert(Array.isArray(strandSort([25, 32, 12, 7, 20])));
```

`strandSort([25, 32, 12, 7, 20])`は`[7, 12, 20, 25, 32]`を返すべきです。

```js
assert.deepEqual(strandSort([25, 32, 12, 7, 20]), [7, 12, 20, 25, 32]);
```

`strandSort([38, 45, 35, 8, 13])`は`[8, 13, 35, 38, 45]`を返すべきです。

```js
assert.deepEqual(strandSort([38, 45, 35, 8, 13]), [8, 13, 35, 38, 45]);
```

`strandSort([43, 36, 20, 34, 24])`は`[20, 24, 34, 36, 43]`を返すべきです。

```js
assert.deepEqual(strandSort([43, 36, 20, 34, 24]), [20, 24, 34, 36, 43]);
```

`strandSort([12, 33, 26, 18, 1, 16, 38])`は`[1, 12, 16, 18, 26, 33, 38]`を返すべきです。

```js
assert.deepEqual(strandSort([12, 33, 26, 18, 1, 16, 38]), [
  1,
  12,
  16,
  18,
  26,
  33,
  38
]);
```

`strandSort([3, 39, 48, 16, 1, 4, 29])`は`[1, 3, 4, 16, 29, 39, 48]`を返すべきです。

```js
assert.deepEqual(strandSort([3, 39, 48, 16, 1, 4, 29]), [
  1,
  3,
  4,
  16,
  29,
  39,
  48
]);
```

# --seed--

## --seed-contents--

```js
function strandSort(list) {

}
```

# --solutions--

```js
function strandSort(list) {
  function merge(left, right) {
    var result = [];
    while (left.length != 0 && right.length != 0) {
      if (left[0] <= right[0]) result.push(left.shift());
      else result.push(right.shift());
    }
    result.push.apply(result, left);
    result.push.apply(result, right);
    return result;
  }

  if (list.length <= 1) return list;
  var result = [];
  while (list.length > 0) {
    var sorted = [];
    sorted.push(list.shift());
    var len = list.length;
    for (var i = 1; i < len; i++) {
      var elem = list[i];
      if (sorted[i - 1] <= elem) {
        sorted.push(elem);
        sorted.splice(i, 1);
      }
    }

    result = merge(result, sorted);
  }
  return result;
}
```
