---
id: 5a23c84252665b21eecc8017
title: Soundex
challengeType: 1
forumTopicId: 302320
dashedName: soundex
---

# --description--

**Soundexアルゴリズム**は単語の*意図*に着目します。似た発音の単語を表す表現を作成します。

これは<em>名前</em>や<em>住所</em>の検索に使われます。つまり、<em>名前</em>を入力した人は、<em>名前</em>の綴りを正す代わりに発音にフォーカスできます。

例えば：

初めて`Quenci`という名前を聞いて、綴りを間違えた場合でも、**Soundex**コードは`Q520`になります。

次に`Quincy`という名前を正しく綴ったとしても、同じコード`Q520`が得られます。これは、複数の名前の発音を同じ<em>人</em>に紐づけられることを意味し、すべての綴りを追加する必要がありません。

ルールは以下の通りです：

<ul>
  <li>母音（A、E、I、O、U）が同じSoundexコードの2つの子音を区切る場合、母音の右側の子音がコード化されます。例えばTymczakはT-522とコード化されます（T、Mに対して5、Cに対して2、Zは無視（上記の「隣接」ルール参照）、Kに対して2）。母音「A」がZとKを区切っているため、Kがコード化されます。</li>
  <li>「H」または「W」が同じSoundexコードの2つの子音を区切る場合、母音の右側の子音はコード化されません。例：AshcraftはA-261とコード化されます（A、Sに対して2、Cは無視、Rに対して6、Fに対して1）。A-226とはなりません。</li>
</ul>

# --instructions--

文字列をパラメータとして受け取り、エンコードされた文字列を返す関数を書いてください。

# --hints--

`soundex`は関数であるべきです。

```js
assert(typeof soundex == 'function');
```

`soundex("Soundex")`は文字列を返すべきです。

```js
assert(typeof soundex('Soundex') == 'string');
```

`soundex("Soundex")`は`"S532"`を返すべきです。

```js
assert.equal(soundex('Soundex'), 'S532');
```

`soundex("Example")`は`"E251"`を返すべきです。

```js
assert.equal(soundex('Example'), 'E251');
```

`soundex("Sownteks")`は`"S532"`を返すべきです。

```js
assert.equal(soundex('Sownteks'), 'S532');
```

`soundex("Ekzampul")`は`"E251"`を返すべきです。

```js
assert.equal(soundex('Ekzampul'), 'E251');
```

`soundex("Euler")`は`"E460"`を返すべきです。

```js
assert.equal(soundex('Euler'), 'E460');
```

`soundex("Gauss")`は`"G200"`を返すべきです。

```js
assert.equal(soundex('Gauss'), 'G200');
```

`soundex("Hilbert")`は`"H416"`を返すべきです。

```js
assert.equal(soundex('Hilbert'), 'H416');
```

`soundex("Knuth")`は`"K530"`を返すべきです。

```js
assert.equal(soundex('Knuth'), 'K530');
```

`soundex("Lloyd")`は`"L300"`を返すべきです。

```js
assert.equal(soundex('Lloyd'), 'L300');
```

`soundex("Lukasiewicz")`は`"L222"`を返すべきです。

```js
assert.equal(soundex('Lukasiewicz'), 'L222');
```

# --seed--

## --seed-contents--

```js
function soundex(s) {

}
```

# --solutions--

```js
function soundex(s) {
  var a = s.toLowerCase().split('');
  var f = a.shift(),
    r = '',
    codes = {
      a: '',
      e: '',
      i: '',
      o: '',
      u: '',
      b: 1,
      f: 1,
      p: 1,
      v: 1,
      c: 2,
      g: 2,
      j: 2,
      k: 2,
      q: 2,
      s: 2,
      x: 2,
      z: 2,
      d: 3,
      t: 3,
      l: 4,
      m: 5,
      n: 5,
      r: 6
    };
  r =
    f +
    a
      .map(function(v, i, a) {
        return codes[v];
      })
      .filter(function(v, i, a) {
        return i === 0 ? v !== codes[f] : v !== a[i - 1];
      })
      .join('');

  return (r + '000').slice(0, 4).toUpperCase();
}
```
