---
id: 5e4ce2f5ac708cc68c1df261
title: 線形合同法ジェネレーター
challengeType: 1
forumTopicId: 385266
dashedName: linear-congruential-generator
---

# --description--

線形合同法ジェネレーター（LCG）は、区分的な線形方程式で計算される疑似乱数のシーケンスを生成する<em>アルゴリズム</em>です。すべての線形合同法ジェネレーターは次の公式を使います：

$$r_{n + 1} = (a \times r_n + c) \bmod m$$

ここで：

<ul>
<li>$ r_0 $ はシードです。</li>
<li>$r_1$, $r_2$, $r_3$, ... は乱数です。</li>
<li>$a$, $c$, $m$ は定数です。</li>
</ul>

$a$, $c$、$m$ の値を慎重に選べば、$0$ から $m - 1$ までの整数の一様分布を生成できます。

<abbr title="linear congruential generator">LCG</abbr> の数値は品質が低いです。$r_n$ と $r_{n + 1}$ は真の乱数のように独立していません。$r_n$ を知っている人は $r_{n + 1}$ を予測できるため、<abbr title="linear congruential generator">LCG</abbr> は暗号学的に安全ではありません。とはいえ、<abbr title="linear congruential generator">LCG</abbr> はミラーラビン素数判定やFreeCellの配り方のような単純なタスクには十分使えます。<abbr title="linear congruential generator">LCG</abbr> の利点の一つは、同じ $r_0$ から数列を簡単に再現できることです。また、公式が非常に単純なので、異なるプログラミング言語でも同じシーケンスを再現できます。

# --instructions--

$r_0,a,c,m,n$ をパラメータとして受け取り、$r_n$ を返す関数を書いてください。

# --hints--

`linearCongGenerator` は関数であるべきです。

```js
assert(typeof linearCongGenerator == 'function');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` は数値を返すべきです。

```js
assert(typeof linearCongGenerator(324, 1145, 177, 2148, 3) == 'number');
```

`linearCongGenerator(324, 1145, 177, 2148, 3)` は `855` を返すべきです。

```js
assert.equal(linearCongGenerator(324, 1145, 177, 2148, 3), 855);
```

`linearCongGenerator(234, 11245, 145, 83648, 4)` は `1110` を返すべきです。

```js
assert.equal(linearCongGenerator(234, 11245, 145, 83648, 4), 1110);
```

`linearCongGenerator(85, 11, 1234, 214748, 5)` は `62217` を返すべきです。

```js
assert.equal(linearCongGenerator(85, 11, 1234, 214748, 5), 62217);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 1)` は `12345` を返すべきです。

```js
assert.equal(linearCongGenerator(0, 1103515245, 12345, 2147483648, 1), 12345);
```

`linearCongGenerator(0, 1103515245, 12345, 2147483648, 2)` は `1406932606` を返すべきです。

```js
assert.equal(
  linearCongGenerator(0, 1103515245, 12345, 2147483648, 2),
  1406932606
);
```

# --seed--

## --seed-contents--

```js
function linearCongGenerator(r0, a, c, m, n) {

}
```

# --solutions--

```js
function linearCongGenerator(r0, a, c, m, n) {
    for (let i = 0; i < n; i++) {
        r0 = (a * r0 + c) % m;
    }
    return r0;
}
```
