---
id: 5e6decd8ec8d7db960950d1c
title: LU分解
challengeType: 1
forumTopicId: 385280
dashedName: lu-decomposition
---

# --description--

任意の正方行列 $A$ は、下三角行列 $L$ と上三角行列 $U$ の積に分解できます。これをLU分解と呼びます。

$A = LU$

これはガウスの消去法の修正版です。

コレスキー分解は対称かつ正定値の行列にのみ適用できますが、より一般的なLU分解は任意の正方行列に対して適用できます。

$L$ と $U$ を計算するためのいくつかのアルゴリズムがあります。

3x3の例で*クルートのアルゴリズム*を導出するには、次の連立方程式を解く必要があります。

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix}= \\begin{pmatrix} l\_{11} & 0 & 0 \\\\ l\_{21} & l\_{22} & 0 \\\\ l\_{31} & l\_{32} & l\_{33}\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = LU\\end{align}

この時点で、12個の未知数に対して9個の方程式を解く必要があります。連立方程式を一意に解けるようにするために、通常は $L$ の対角要素を1に設定します。

$l\_{11}=1$

$l\_{22}=1$

$l\_{33}=1$

こうすることで、9個の未知数と9個の方程式の解ける連立方程式になります。

\\begin{align}A = \\begin{pmatrix} a\_{11} & a\_{12} & a\_{13}\\\\ a\_{21} & a\_{22} & a\_{23}\\\\ a\_{31} & a\_{32} & a\_{33}\\\\ \\end{pmatrix} = \\begin{pmatrix} 1 & 0 & 0 \\\\ l\_{21} & 1 & 0 \\\\ l\_{31} & l\_{32} & 1\\\\ \\end{pmatrix} \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ 0 & u\_{22} & u\_{23} \\\\ 0 & 0 & u\_{33} \\end{pmatrix} = \\begin{pmatrix} u\_{11} & u\_{12} & u\_{13} \\\\ u\_{11}l\_{21} & u\_{12}l\_{21}+u\_{22} & u\_{13}l\_{21}+u\_{23} \\\\ u\_{11}l\_{31} & u\_{12}l\_{31}+u\_{22}l\_{32} & u\_{13}l\_{31} + u\_{23}l\_{32}+u\_{33} \\end{pmatrix} = LU\\end{align}

他の $l$ と $u$ を解くと、次の方程式が得られます。

$u\_{11}=a\_{11}$

$u\_{12}=a\_{12}$

$u\_{13}=a\_{13}$

$u\_{22}=a\_{22} - u\_{12}l\_{21}$

$u\_{23}=a\_{23} - u\_{13}l\_{21}$

$u\_{33}=a\_{33} - (u\_{13}l\_{31} + u\_{23}l\_{32})$

そして $l$ については：

$l\_{21}=\\frac{1}{u\_{11}} a\_{21}$

$l\_{31}=\\frac{1}{u\_{11}} a\_{31}$

$l\_{32}=\\frac{1}{u\_{22}} (a\_{32} - u\_{12}l\_{31})$

計算パターンが見えてきて、次の式で表現できます。まず $U$ について

$u\_{ij} = a\_{ij} - \\sum\_{k=1}^{i-1} u\_{kj}l\_{ik}$

次に $L$ について

$l\_{ij} = \\frac{1}{u\_{jj}} (a\_{ij} - \\sum\_{k=1}^{j-1} u\_{kj}l\_{ik})$

2番目の式から、対角要素（ピボット） $u\_{jj}$ で割る必要があるため、$u\_{jj}$ が0または非常に小さい場合に数値的不安定性が生じる問題がわかります。

この問題の解決策は、*ピボット選択*で、LU分解の前に $A$ の行を並べ替えて、各列の最大要素が $A$ の対角に来るようにします。行の並べ替えは、$A$ に置換行列 $P$ を掛けることを意味します。

$PA \\Rightarrow A'$

例：

\\begin{align} \\begin{pmatrix} 0 & 1 \\\\ 1 & 0 \\end{pmatrix} \\begin{pmatrix} 1 & 4 \\\\ 2 & 3 \\end{pmatrix} \\Rightarrow \\begin{pmatrix} 2 & 3 \\\\ 1 & 4 \\end{pmatrix} \\end{align}

その後、並べ替えた行列に対して分解アルゴリズムを適用し、

$PA = LU$

となります。

# --instructions--

正方の $n \times n$ 行列 $A$ を受け取り、下三角行列 $L$、上三角行列 $U$、および置換行列 $P$ を返すルーチンを実装してください。これにより上記の等式が成り立ちます。返り値は `[L, U, P]` の形式にしてください。

# --before-each--

```js
const matrix1 = [
  [1, 3, 5],
  [2, 4, 7],
  [1, 1, 0]
];
const expected1 = [
  [
    [1, 0, 0],
    [0.5, 1, 0],
    [0.5, -1, 1]
  ],
  [
    [2, 4, 7],
    [0, 1, 1.5],
    [0, 0, -2]
  ],
  [
    [0, 1, 0],
    [1, 0, 0],
    [0, 0, 1]
  ]
];

const matrix2 = [
  [11, 9, 24, 2],
  [1, 5, 2, 6],
  [3, 17, 18, 1],
  [2, 5, 7, 1]
];
const expected2 = [
  [
    [1, 0, 0, 0],
    [0.2727272727272727, 1, 0, 0],
    [0.09090909090909091, 0.2875, 1, 0],
    [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]
  ],
  [
    [11, 9, 24, 2],
    [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546],
    [0, 0, -3.4749999999999996, 5.6875],
    [0, 0, 0, 0.510791366906476]
  ],
  [
    [1, 0, 0, 0],
    [0, 0, 1, 0],
    [0, 1, 0, 0],
    [0, 0, 0, 1]
  ]
];

const matrix3 = [
  [1, 1, 1],
  [4, 3, -1],
  [3, 5, 3]
];
const expected3 = [
  [
    [1, 0, 0],
    [0.75, 1, 0],
    [0.25, 0.09090909090909091, 1]
  ],
  [
    [4, 3, -1],
    [0, 2.75, 3.75],
    [0, 0, 0.9090909090909091]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];

const matrix4 = [
  [1, -2, 3],
  [2, -5, 12],
  [0, 2, -10]
];
const expected4 = [
  [
    [1, 0, 0],
    [0, 1, 0],
    [0.5, 0.25, 1]
  ],
  [
    [2, -5, 12],
    [0, 2, -10],
    [0, 0, -0.5]
  ],
  [
    [0, 1, 0],
    [0, 0, 1],
    [1, 0, 0]
  ]
];
```

# --hints--

`luDecomposition` は関数であるべきです。

```js
assert(typeof luDecomposition == 'function');
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` は配列を返すべきです。

```js
assert(Array.isArray(luDecomposition(matrix1)));
```

`luDecomposition([[1, 3, 5], [2, 4, 7], [1, 1, 0]])` は `[[[1, 0, 0], [0.5, 1, 0], [0.5, -1, 1]], [[2, 4, 7], [0, 1, 1.5], [0, 0, -2]], [[0, 1, 0], [1, 0, 0], [0, 0, 1]]]` を返すべきです。

```js
assert.deepEqual(luDecomposition(matrix1), expected1);
```

`luDecomposition([[11, 9, 24, 2], [1, 5, 2, 6], [3, 17, 18, 1], [2, 5, 7, 1]])` は `[[[1, 0, 0, 0], [0.2727272727272727, 1, 0, 0], [0.09090909090909091, 0.2875, 1, 0], [0.18181818181818182, 0.23124999999999996, 0.0035971223021580693, 1]], [[11, 9, 24, 2], [0, 14.545454545454547, 11.454545454545455, 0.4545454545454546], [0, 0, -3.4749999999999996, 5.6875], [0, 0, 0, 0.510791366906476]], [[1, 0, 0, 0], [0, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 1]]]` を返すべきです。

```js
assert.deepEqual(luDecomposition(matrix2), expected2);
```

`luDecomposition([[1, 1, 1], [4, 3, -1], [3, 5, 3]])` は `[[[1, 0, 0], [0.75, 1, 0], [0.25, 0.09090909090909091, 1]], [[4, 3, -1], [0, 2.75, 3.75], [0, 0, 0.9090909090909091]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]` を返すべきです。

```js
assert.deepEqual(luDecomposition(matrix3), expected3);
```

`luDecomposition([[1, -2, 3], [2, -5, 12], [0, 2, -10]])` は `[[[1, 0, 0], [0, 1, 0], [0.5, 0.25, 1]], [[2, -5, 12], [0, 2, -10], [0, 0, -0.5]], [[0, 1, 0], [0, 0, 1], [1, 0, 0]]]` を返すべきです。

```js
assert.deepEqual(luDecomposition(matrix4), expected4);
```

# --seed--

## --seed-contents--

```js
function luDecomposition(A) {

}
```

# --solutions--

```js
function luDecomposition(A) {

    function dotProduct(a, b) {
        var sum = 0;
        for (var i = 0; i < a.length; i++)
            sum += a[i] * b[i]
        return sum;
    }

    function matrixMul(A, B) {
        var result = new Array(A.length);
        for (var i = 0; i < A.length; i++)
            result[i] = new Array(B[0].length)
        var aux = new Array(B.length);

        for (var j = 0; j < B[0].length; j++) {

            for (var k = 0; k < B.length; k++)
                aux[k] = B[k][j];

            for (var i = 0; i < A.length; i++)
                result[i][j] = dotProduct(A[i], aux);
        }
        return result;
    }

    function pivotize(m) {
        var n = m.length;
        var id = new Array(n);
        for (var i = 0; i < n; i++) {
            id[i] = new Array(n);
            id[i].fill(0)
            id[i][i] = 1;
        }

        for (var i = 0; i < n; i++) {
            var maxm = m[i][i];
            var row = i;
            for (var j = i; j < n; j++)
                if (m[j][i] > maxm) {
                    maxm = m[j][i];
                    row = j;
                }

            if (i != row) {
                var tmp = id[i];
                id[i] = id[row];
                id[row] = tmp;
            }
        }
        return id;
    }

    var n = A.length;
    var L = new Array(n);
    for (var i = 0; i < n; i++) { L[i] = new Array(n); L[i].fill(0) }
    var U = new Array(n);
    for (var i = 0; i < n; i++) { U[i] = new Array(n); U[i].fill(0) }
    var P = pivotize(A);
    var A2 = matrixMul(P, A);

    for (var j = 0; j < n; j++) {
        L[j][j] = 1;
        for (var i = 0; i < j + 1; i++) {
            var s1 = 0;
            for (var k = 0; k < i; k++)
                s1 += U[k][j] * L[i][k];
            U[i][j] = A2[i][j] - s1;
        }
        for (var i = j; i < n; i++) {
            var s2 = 0;
            for (var k = 0; k < j; k++)
                s2 += U[k][j] * L[i][k];
            L[i][j] = (A2[i][j] - s2) / U[j][j];
        }
    }
    return [L, U, P];
}
```
