---
id: 5e6dd14797f5ce267c2f19d0
title: ルック・アンド・セイ・シーケンス
challengeType: 1
forumTopicId: 385277
dashedName: look-and-say-sequence
---

# --description--

ルック・アンド・セイ・シーケンスは再帰的に定義された数列です。

シーケンスの定義

<ul><li>10進数の数字を取ります</li>
<li>数字を見て、同じ数字が連続している部分を視覚的にグループ化します。</li>
<li>左から右へ、グループごとに「何個のその数字があるか」を言います。その後にグループ化された数字を続けます。</li></ul><span> これが次のシーケンスの数字になります。</span>

例：

<ul><li>数字1から始めると、1が<span>1つ</span>あるので11になります</li>
<li>11から始めると、1が<span>2つ</span>あります。つまり21です</li>
<li>21から始めると、2が<span>1つ</span>、次に1が<span>1つ</span>あります。つまり(12)(11)で1211になります</li>
<li>1211から始めると、1が<span>1つ</span>、2が<span>1つ</span>、次に1が<span>2つ</span>あります。つまり(11)(12)(21)で111221になります</li></ul>

# --instructions--

文字列をパラメータとして受け取り、それを処理して結果の文字列を返す関数を書いてください。

# --hints--

`lookAndSay` は関数であるべきです。

```js
assert(typeof lookAndSay == 'function');
```

`lookAndSay("1")` は文字列を返すべきです。

```js
assert(typeof lookAndSay('1') == 'string');
```

`lookAndSay("1")` は `"11"` を返すべきです。

```js
assert.equal(lookAndSay('1'), '11');
```

`lookAndSay("11")` は `"21"` を返すべきです。

```js
assert.equal(lookAndSay('11'), '21');
```

`lookAndSay("21")` は `"1211"` を返すべきです。

```js
assert.equal(lookAndSay('21'), '1211');
```

`lookAndSay("1211")` は `"111221"` を返すべきです。

```js
assert.equal(lookAndSay('1211'), '111221');
```

`lookAndSay("3542")` は `"13151412"` を返すべきです。

```js
assert.equal(lookAndSay('3542'), '13151412');
```

# --seed--

## --seed-contents--

```js
function lookAndSay(str) {

}
```

# --solutions--

```js
function lookAndSay(str) {
    return str.replace(/(.)\1*/g, function(seq, p1) {
      return seq.length.toString() + p1;
    });
}
```
