---
id: 5e6dd1278e6ca105cde40ea9
title: 最長共通部分列
challengeType: 1
forumTopicId: 385271
dashedName: longest-common-subsequence
---

# --description--

**最長共通部分列**（または**LCS**）とは、グループAとBの要素のうち、両方のグループに共通し、かつそれぞれのグループ内で同じ順序で並んでいる最も長い要素の集まりのことです。例えば、シーケンス`1234`と`1224533324`のLCSは`1234`です。

<u>1234</u>

<u>12</u>245<u>3</u>332<u>4</u>

文字列の例として、シーケンス`thisisatest`と`testing123testing`を考えます。LCSは`tsitest`となります。

<u>t</u>hi<u>si</u>sa<u>test</u>

<u>t</u>e<u>s</u>t<u>i</u>ng123<u>test</u>ing.

コードは文字列のみを扱う必要があります。

# --instructions--

大文字・小文字を区別する関数を書いて、2つの文字列のLCSを返してください。複数のLCSを表示する必要はありません。

# --hints--

`lcs`は関数であるべきです。

```js
assert(typeof lcs == 'function');
```

`lcs("thisisatest", "testing123testing")`は文字列を返すべきです。

```js
assert(typeof lcs('thisisatest', 'testing123testing') == 'string');
```

`lcs("thisisatest", "testing123testing")`は`"tsitest"`を返すべきです。

```js
assert.equal(lcs('thisisatest', 'testing123testing'), 'tsitest');
```

`lcs("ABCDGH", "AEDFHR")`は`"ADH"`を返すべきです。

```js
assert.equal(lcs('ABCDGH', 'AEDFHR'), 'ADH');
```

`lcs("AGGTAB", "GXTXAYB")`は`"GTAB"`を返すべきです。

```js
assert.equal(lcs('AGGTAB', 'GXTXAYB'), 'GTAB');
```

`lcs("BDACDB", "BDCB")`は`"BDCB"`を返すべきです。

```js
assert.equal(lcs('BDACDB', 'BDCB'), 'BDCB');
```

`lcs("ABAZDC", "BACBAD")`は`"ABAD"`を返すべきです。

```js
assert.equal(lcs('ABAZDC', 'BACBAD'), 'ABAD');
```

# --seed--

## --seed-contents--

```js
function lcs(a, b) {

}
```

# --solutions--

```js
function lcs(a, b) {
  var aSub = a.substring(0, a.length - 1);
  var bSub = b.substring(0, b.length - 1);

  if (a.length === 0 || b.length === 0) {
    return '';
  } else if (a.charAt(a.length - 1) === b.charAt(b.length - 1)) {
    return lcs(aSub, bSub) + a.charAt(a.length - 1);
  } else {
    var x = lcs(a, bSub);
    var y = lcs(aSub, b);
    return (x.length > y.length) ? x : y;
  }
}
```
