---
id: 5e6dd139859c290b6ab80292
title: 最長増加部分列
challengeType: 1
forumTopicId: 385272
dashedName: longest-increasing-subsequence
---

# --description--

最長増加部分列問題とは、与えられたシーケンスの中から、要素が昇順に並んでいる部分列を見つけ、その部分列ができるだけ長くなるようにする問題です。例を示します。

次の配列の場合：

```js
const array = [3, 10, 2, 1, 20];
```

最長増加部分列は：

$\\{3, 10, 20\\}$

# --instructions--

数値の配列をパラメータとして受け取り、最長増加部分列を返す関数を書いてください。

すべての配列には最長増加部分列が必ず存在します。

# --hints--

`findSequence` は関数である必要があります。

```js
assert(typeof findSequence == 'function');
```

`findSequence([3, 10, 2, 1, 20])` は配列を返す必要があります。

```js
assert(Array.isArray(findSequence([3, 10, 2, 1, 20])));
```

`findSequence([3, 10, 2, 1, 20])` は `[3, 10, 20]` を返す必要があります。

```js
assert.deepEqual(findSequence([3, 10, 2, 1, 20]), [3, 10, 20]);
```

`findSequence([2, 7, 3, 5, 8])` は `[2, 3, 5, 8]` を返す必要があります。

```js
assert.deepEqual(findSequence([2, 7, 3, 5, 8]), [2, 3, 5, 8]);
```

`findSequence([2, 6, 4, 5, 1])` は `[2, 4, 5]` を返す必要があります。

```js
assert.deepEqual(findSequence([2, 6, 4, 5, 1]), [2, 4, 5]);
```

`findSequence([10, 22, 9, 33, 21, 50, 60, 80])` は `[10, 22, 33, 50, 60, 80]` を返す必要があります。

```js
assert.deepEqual(findSequence([10, 22, 9, 33, 21, 50, 60, 80]), [
  10,
  22,
  33,
  50,
  60,
  80
]);
```

`findSequence([0, 12, 2, 10, 6, 14, 1, 9, 5, 13, 3, 11, 7, 15])` は `[0, 2, 6, 9, 11, 15` を返す必要があります。

```js
assert.deepEqual(
  findSequence([0, 12, 2, 10, 6, 14, 1, 9, 5, 13, 3, 11, 7, 15]),
  [0, 2, 6, 9, 11, 15]
);
```

# --seed--

## --seed-contents--

```js
function findSequence(input) {

}
```

# --solutions--

```js
function findSequence(input) {
    var len = input.length;
    var result = []
    for (var i = 0; i < len; i++) result.push(1)

    for (var i = 0; i < len; i++)
        for (var j = i - 1; j >= 0; j--)
            if (input[i] > input[j] && result[j] >= result[i])
                result[i] = result[j] + 1;

    var maxValue = Math.max.apply(null, result);
    var maxIndex = result.indexOf(Math.max.apply(Math, result));
    var output = [];
    output.push(input[maxIndex]);
    for (var i = maxIndex; i >= 0; i--) {
        if (maxValue == 0) break;
        if (input[maxIndex] > input[i] && result[i] == maxValue - 1) {
            output.push(input[i]);
            maxValue--;
        }
    }
    output.reverse();
    return output;
}
```
