---
id: 5e94a54cc7b022105bf0fd2c
title: 単語の出現頻度
challengeType: 1
forumTopicId: 393913
dashedName: word-frequency
---

# --description--

文字列と整数 n が与えられたとき、ファイル内で最も頻繁に出現する n 個の単語とその出現回数を、出現頻度の高い順に返してください。

# --instructions--

各単語の出現回数を数え、出現頻度の高い順に n 個の単語とその出現回数を返す関数を書いてください。

関数は 2 次元配列を返す必要があり、各要素は次の形式です：`[word, freq]`。`word` は単語の小文字バージョン、`freq` は出現回数を表します。

文字列が与えられなかった場合は、空の配列を返してください。

大文字・小文字は区別しません。たとえば "Hello" と "hello" は同じ単語として扱います。

アンダースコア、ダッシュ、アポストロフィ、カンマなどの特殊文字を含む単語は、それぞれ別の単語として扱ってかまいません。

たとえば、文字列 "Hello hello goodbye" が与えられた場合、関数は `[['hello', 2], ['goodbye', 1]]` を返すべきです。

# --hints--

`wordFrequency` は関数であるべきです。

```js
assert.isFunction(wordFrequency);
```

`wordFrequency` は配列を返すべきです。

```js
assert.isArray(wordFrequency('test'));
```

`wordFrequency("Hello hello world", 2)` は `[['hello', 2], ['world', 1]]` を返すべきです。

```js
const example_1 = 'Hello hello world';
const example_1_solution = [['hello', 2], ['world', 1]];

assert.deepEqual(wordFrequency(example_1, 2), example_1_solution);
```

`wordFrequency("The quick brown fox jumped over the lazy dog", 1)` は `[['the', 2]]` を返すべきです。

```js
const example_2 = 'The quick brown fox jumped over the lazy dog';
const example_2_solution = [['the', 2]];

assert.deepEqual(wordFrequency(example_2, 1), example_2_solution);
```

`wordFrequency("Opensource opensource open-source open source", 1)` は `[['opensource', 2]]` を返すべきです。

```js
const example_3 = 'Opensource opensource open-source open source';
const example_3_solution = [['opensource', 2]];

assert.deepEqual(wordFrequency(example_3, 1), example_3_solution);
```

`wordFrequency("Apple App apply aPP aPPlE", 3)` は `[['app', 2], ['apple', 2], ['apply', 1]]` または `[['apple', 2], ['app', 2], ['apply', 1]]` を返すべきです。

```js
const example_4 = 'Apple App apply aPP aPPlE';
const example_4_solution = [["app", 2],["apple", 2],["apply", 1]];

const arr = wordFrequency(example_4, 3);
assert.includeDeepMembers(arr, example_4_solution);
```

`wordFrequency("c d a d c a b d d c", 4)` は `[['d', 4], ['c', 3], ['a', 2], ['b', 1]]` を返すべきです。

```js
const example_5 = 'c d a d c a b d d c';
const example_5_solution = [['d', 4], ['c', 3], ['a', 2], ['b', 1]];

assert.deepEqual(wordFrequency(example_5, 4), example_5_solution);
```

`wordFrequency("", 5)` は `[]` を返すべきです。

```js
const example_6 = '';
const example_6_solution = [];

assert.deepEqual(wordFrequency(example_6, 5), example_6_solution);
```

# --seed--

## --seed-contents--

```js
function wordFrequency(txt, n) {

}
```

# --solutions--

```js
function wordFrequency(txt, n) {
  var words = txt.split(/\s+/);
  var wordCount = {};
  words.forEach(word => {
    if (word == '') {
      return;
    }
    const lowerWord = word.toLowerCase();
    wordCount[lowerWord] =
      lowerWord in wordCount ? wordCount[lowerWord] + 1 : 1;
  });

  var wordsArray = [];
  for (let [word, count] of Object.entries(wordCount)) {
    wordsArray.push([word, count]);
  }

  wordsArray.sort((a, b) => {
    if (a[1] !== b[1]) {
      return b[1] - a[1];
    } else if (a[0] !== b[0]) {
      return a[0] < b[0] ? -1 : 1;
    }
    return 0;
  });
  return wordsArray.slice(0, n);
}
```
