---
id: 5e9ddb06ec35240f39657419
title: FizzBuzz
challengeType: 1
forumTopicId: 385370
dashedName: fizzbuzz
---

# --description--

1から100まで（100を含む）の整数の配列を生成するプログラムを書いてください。ただし：

<ul>
    <li>3の倍数の場合は、数値の代わりに<code>"Fizz"</code>を配列に追加してください</li>
    <li>5の倍数の場合は、数値の代わりに<code>"Buzz"</code>を配列に追加してください</li>
    <li>3と5の両方の倍数の場合は、数値の代わりに<code>"FizzBuzz"</code>を配列に追加してください</li>
</ul>

# --instructions--

上記のルールに基づいた結果を含む配列を返すプログラムにしてください。

# --hints--

`fizzBuzz` は関数である必要があります。

```js
assert(typeof fizzBuzz == 'function');
```

`fizzBuzz()` は配列を返す必要があります。

```js
assert(Array.isArray(fizzBuzz()) == true);
```

3で割り切れる数は `"Fizz"` を返す必要があります。

```js
assert.equal(fizzBuzz()[2], 'Fizz');
```

5で割り切れる数は `"Buzz"` を返す必要があります。

```js
assert.equal(fizzBuzz()[99], 'Buzz');
```

3と5の両方で割り切れる数は `"FizzBuzz"` を返す必要があります。

```js
assert.equal(fizzBuzz()[89], 'FizzBuzz');
```

3でも5でも割り切れない数は、その数自体を返す必要があります。

```js
assert.equal(fizzBuzz()[12], 13);
```

# --seed--

## --seed-contents--

```js
function fizzBuzz() {

}
```

# --solutions--

```js
function fizzBuzz() {
    let res=[];
    for (let i =1; i < 101; i++) {
        if (i % 3 === 0  && i % 5 === 0) {
            res.push("FizzBuzz");
        }
        else if (i % 3 === 0) {
            res.push("Fizz");
        }
        else if (i % 5 === 0) {
            res.push("Buzz");
        } 
        else {
            res.push(i);
        }
    }
    return res;
}
```
