---
id: 5ea28156e79528a9ab248f27
title: クレジットカード番号のLuhnテスト
challengeType: 1
forumTopicId: 385284
dashedName: luhn-test-of-credit-card-numbers
---

# --description--

Luhnテストは、一部のクレジットカード会社が有効なクレジットカード番号を、ランダムな数字の組み合わせと区別するために使っています。

Luhnテストで検証できるクレジットカード番号を使っている会社の番号は、以下のテストに合格します。

<ol>
  <li>番号の数字の順序を逆にします。</li>
  <li>逆にした数字のうち、1番目、3番目、…といった奇数番目の数字を取り出し、それらを合計して部分和s1を作ります。</li>
  <li>逆にした数字のうち、2番目、4番目、…といった偶数番目の数字については：</li>
    <ol>
      <li>各数字を2倍し、もし答えが9より大きければその数字の各桁を足して、偶数番目の数字の部分和を作ります。</li>
      <li>偶数番目の数字の部分和を合計してs2を作ります。</li>
    </ol>
  <li>s1 + s2の合計が0で終わる場合、元の番号はLuhnテストで検証された有効なクレジットカード番号の形式です。</li>
</ol>

例えば、試験番号が49927398716の場合：

```md
Reverse the digits:
  61789372994
Sum the odd digits:
  6 + 7 + 9 + 7 + 9 + 4 = 42 = s1
The even digits:
    1, 8, 3, 2, 9
  Two times each even digit:
    2, 16, 6, 4, 18
  Sum the digits of each multiplication:
    2, 7, 6, 4, 9
  Sum the last:
    2 + 7 + 6 + 4 + 9 = 28 = s2

s1 + s2 = 70 which ends in zero which means that 49927398716 passes the Luhn test.
```

# --instructions--

Luhnテストで番号を検証する関数を書いてください。有効な番号ならtrueを返し、そうでなければfalseを返してください。

# --hints--

`luhnTest` は関数であるべきです。

```js
assert(typeof luhnTest === 'function');
```

`luhnTest("4111111111111111")` はブール値を返すべきです。

```js
assert(typeof luhnTest('4111111111111111') === 'boolean');
```

`luhnTest("4111111111111111")` は `true` を返すべきです。

```js
assert.equal(luhnTest('4111111111111111'), true);
```

`luhnTest("4111111111111112")` は `false` を返すべきです。

```js
assert.equal(luhnTest('4111111111111112'), false);
```

`luhnTest("49927398716")` は `true` を返すべきです。

```js
assert.equal(luhnTest('49927398716'), true);
```

`luhnTest("49927398717")` は `false` を返すべきです。

```js
assert.equal(luhnTest('49927398717'), false);
```

`luhnTest("1234567812345678")` は `false` を返すべきです。

```js
assert.equal(luhnTest('1234567812345678'), false);
```

`luhnTest("1234567812345670")` は `true` を返すべきです。

```js
assert.equal(luhnTest('1234567812345670'), true);
```

# --seed--

## --seed-contents--

```js
function luhnTest(str) {

}
```

# --solutions--

```js
function luhnTest(str) {
  var luhnArr = [0, 2, 4, 6, 8, 1, 3, 5, 7, 9];
  var counter = 0;
  var incNum;
  var odd = false;
  var temp = String(str).replace(/[^\d]/g, '');
  if (temp.length == 0) return false;
  for (var i = temp.length - 1; i >= 0; --i) {
    incNum = parseInt(temp.charAt(i), 10);
    counter += (odd = !odd) ? incNum : luhnArr[incNum];
  }
  return counter % 10 == 0;
}
```
