---
id: 5eaf48389ee512d4d103684b
title: 自己記述数
challengeType: 1
forumTopicId: 385289
dashedName: self-describing-numbers
---

# --description--

いわゆる「自己記述数」と呼ばれる整数がいくつかあります。

整数が「自己記述数」と呼ばれるのは、桁の位置を0からN-1までラベル付けしたときに、各桁の数字がその数字が数値内に現れる回数と等しい場合です。

例えば、**2020** は4桁の自己記述数です：

<ul>
    <li> 位置0の値は2で、数値内に0が2回現れます； </li>
    <li> 位置1の値は0で、数値内に1は現れません； </li>
    <li> 位置2の値は2で、2が2回現れます； </li>
    <li> 位置3の値は0で、3は0回現れます； </li>
</ul>

100,000,000未満の自己記述数は、1210、2020、21200、3211000、42101000です。

# --instructions--

正の整数をパラメータとして受け取る関数を書いてください。もしその数が自己記述数ならtrueを返してください。そうでなければfalseを返してください。

# --hints--

`isSelfDescribing` は関数である必要があります。

```js
assert(typeof isSelfDescribing == 'function');
```

`isSelfDescribing()` はブール値を返す必要があります。

```js
assert(typeof isSelfDescribing(2020) == 'boolean');
```

`isSelfDescribing(2020)` は `true` を返す必要があります。

```js
assert.equal(isSelfDescribing(2020), true);
```

`isSelfDescribing(3021)` は `false` を返す必要があります。

```js
assert.equal(isSelfDescribing(3021), false);
```

`isSelfDescribing(3211000)` は `true` を返す必要があります。

```js
assert.equal(isSelfDescribing(3211000), true);
```

# --seed--

## --seed-contents--

```js
function isSelfDescribing(n) {

}
```

# --solutions--

```js
function isSelfDescribing(n) {
    let digits = String(n).split("");
    digits = digits.map(function(e) {return parseInt(e)});
    let count = digits.map((x) => {return 0})
    digits.forEach((d) =>{
        if (d >= count.length) {
            return false
        }
        count[d] += 1;
    });

     if (digits === count) {
        return true;
    }
    if (digits.length != count.length) {
        return false;
    }
    
    for (let i=0; i< digits.length; i++){
      if (digits[i] !== count[i]) {
        return false;
      }
    }
    return true;
}
```
