---
id: 5eb3e497b8d6d7f63c5517ea
title: レコードのリストを検索する
challengeType: 1
forumTopicId: 385315
dashedName: search-a-list-of-records
---

# --description--

レコードはエンティティを説明する属性で構成されます。各属性には名前と値があります。たとえば、人は`age`という属性を持ち、その値が25である場合があります。レコードのリストに対して重要な操作は、特定の属性値を持つレコードを見つけることです。

# --instructions--

文字列をパラメータとして受け取る関数を書いてください。その関数は、`list`の中で`name`属性の値が与えられた文字列と一致するアイテムのインデックスを返す必要があります。

# --hints--

`searchCity`は関数である必要があります。

```js
assert(typeof searchCity === 'function');
```

`searchCity("Dar Es Salaam")`は数値を返す必要があります。

```js
assert(typeof searchCity('Dar Es Salaam') === 'number');
```

`searchCity("Dar Es Salaam")`は`6`を返す必要があります。

```js
assert.equal(searchCity('Dar Es Salaam'), 6);
```

`searchCity("Casablanca")`は`9`を返す必要があります。

```js
assert.equal(searchCity('Casablanca'), 9);
```

`searchCity("Cairo")`は`1`を返す必要があります。

```js
assert.equal(searchCity('Cairo'), 1);
```

`searchCity("Mogadishu")`は`4`を返す必要があります。

```js
assert.equal(searchCity('Mogadishu'), 4);
```

`searchCity("Lagos")`は`0`を返す必要があります。

```js
assert.equal(searchCity('Lagos'), 0);
```

# --seed--

## --seed-contents--

```js
function searchCity(name) {

}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```

# --solutions--

```js
function searchCity(name) {
  return list.findIndex(item => item.name === name);
}

const list = [
  { name: 'Lagos', population: 21.0 },
  { name: 'Cairo', population: 15.2 },
  { name: 'Kinshasa-Brazzaville', population: 11.3 },
  { name: 'Greater Johannesburg', population: 7.55 },
  { name: 'Mogadishu', population: 5.85 },
  { name: 'Khartoum-Omdurman', population: 4.98 },
  { name: 'Dar Es Salaam', population: 4.7 },
  { name: 'Alexandria', population: 4.58 },
  { name: 'Abidjan', population: 4.4 },
  { name: 'Casablanca', population: 3.98 }
];
```
