---
id: 5eb3e4b20aa93c437f9e9717
title: 実数の集合
challengeType: 1
forumTopicId: 385322
dashedName: set-of-real-numbers
---

# --description--

すべての実数は非可算な集合 ℝ を形成します。その部分集合の中で比較的単純なのは凸集合で、これは実数 *a* と *b*（*a* ≤ *b*）の間の範囲として表されます。「間」の意味には実際に4つのケースがあり、境界が開いているか閉じているかによって異なります。

<ul>
  <li>[<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> ≤ <i>x</i> かつ <i>x</i> ≤ <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> < <i>x</i> かつ <i>x</i> < <i>b</i> }</li>
  <li>[<i>a</i>, <i>b</i>): {<i>x</i> | <i>a</i> ≤ <i>x</i> かつ <i>x</i> < <i>b</i> }</li>
  <li>(<i>a</i>, <i>b</i>]: {<i>x</i> | <i>a</i> < <i>x</i> かつ <i>x</i> ≤ <i>b</i> }</li>
</ul>

なお、*a* = *b* の場合、4つのうち非空となるのは \[*a*, *a*] のみです。

**課題**

<ul>
  <li>実装ノートの「任意」の定義に従い、任意の実数の集合を表現する方法を考案してください。</li>
  <li>以下の一般的な集合演算のメソッドを提供してください（<i>x</i> は実数、<i>A</i> と <i>B</i> は集合です）：</li>
  <ul>
    <li>
      <i>x</i> ∈ <i>A</i>: <i>x</i> が <i>A</i> の要素かどうか判定する<br>
      例：1 は [1, 2) に含まれますが、2 や 3 などは含まれません。
    </li>
    <li>
      <i>A</i> ∪ <i>B</i>: <i>A</i> と <i>B</i> の和集合、すなわち {<i>x</i> | <i>x</i> ∈ <i>A</i> または <i>x</i> ∈ <i>B</i>}<br>
      例：[0, 2) ∪ (1, 3) = [0, 3); [0, 1) ∪ (2, 3] = そのまま [0, 1) ∪ (2, 3]
    </li>
    <li>
      <i>A</i> ∩ <i>B</i>: <i>A</i> と <i>B</i> の共通部分、すなわち {<i>x</i> | <i>x</i> ∈ <i>A</i> かつ <i>x</i> ∈ <i>B</i>}<br>
      例：[0, 2) ∩ (1, 3) = (1, 2); [0, 1) ∩ (2, 3] = 空集合
    </li>
    <li>
      <i>A</i> - <i>B</i>: <i>A</i> から <i>B</i> を除いた差集合、別表記では <i>A</i> \ <i>B</i>、すなわち {<i>x</i> | <i>x</i> ∈ <i>A</i> かつ <i>x</i> ∉ <i>B</i>}<br>
      例：[0, 2) − (1, 3) = [0, 1]
    </li>
  </ul>
</ul>

# --instructions--

2つのオブジェクト（文字列と配列）をパラメータに取る関数を書いてください。オブジェクトは集合を表し、属性として `low`、`high`、`rangeType` を持ちます。

`rangeType` は `CLOSED`、`BOTH_OPEN`、`LEFT_OPEN`、`RIGHT_OPEN` に対してそれぞれ 0、1、2、3 の値を取ることができます。この情報を使って集合を実装してください。

文字列は集合に対して実行する操作を表し、`"union"`、`"intersect"`、`"subtract"`（差集合） のいずれかです。

操作を実行した後、配列内の値が結果の集合に含まれているかどうかを判定し、それに対応するブール値を配列に格納してください。関数はこの配列を返す必要があります。

# --hints--

`realSet` は関数であるべきです。

```js
assert(typeof realSet == 'function');
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` は配列を返すべきです。

```js
assert(
  Array.isArray(
    realSet(
      { low: 0, high: 1, rangeType: 2 },
      { low: 0, high: 2, rangeType: 3 },
      'union',
      [1, 2, 3]
    )
  )
);
```

`realSet({"low":0, "high":1, "rangeType":2}, {"low":0, "high":2, "rangeType":3}, "union", [1, 2, 3])` は `[true, false, false]` を返すべきです。

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 1, rangeType: 2 },
    { low: 0, high: 2, rangeType: 3 },
    'union',
    [1, 2, 3]
  ),
  [true, false, false]
);
```

`realSet({"low":0, "high":2, "rangeType":3}, {"low":1, "high":2, "rangeType":2}, "intersect", [0, 1, 2])` は `[false, false, false]` を返すべきです。

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 2, rangeType: 3 },
    { low: 1, high: 2, rangeType: 2 },
    'intersect',
    [0, 1, 2]
  ),
  [false, false, false]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":1}, "subtract", [0, 1, 2])` は `[true, true, true]` を返すべきです。

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 1 },
    'subtract',
    [0, 1, 2]
  ),
  [true, true, true]
);
```

`realSet({"low":0, "high":3, "rangeType":3}, {"low":0, "high":1, "rangeType":0}, "subtract", [0, 1, 2])` は `[false, false, true]` を返すべきです。

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 3, rangeType: 3 },
    { low: 0, high: 1, rangeType: 0 },
    'subtract',
    [0, 1, 2]
  ),
  [false, false, true]
);
```

`realSet({"low":0, "high":33, "rangeType":1}, {"low":30, "high":31, "rangeType":0}, "intersect", [30, 31, 32])` は `[true, true, false]` を返すべきです。

```js
assert.deepEqual(
  realSet(
    { low: 0, high: 33, rangeType: 1 },
    { low: 30, high: 31, rangeType: 0 },
    'intersect',
    [30, 31, 32]
  ),
  [true, true, false]
);
```

# --seed--

## --seed-contents--

```js
function realSet(set1, set2, operation, values) {

}
```

# --solutions--

```js
function realSet(set1, set2, operation, values) {
  const RangeType = {
    CLOSED: 0,
    BOTH_OPEN: 1,
    LEFT_OPEN: 2,
    RIGHT_OPEN: 3
  };

  function Predicate(test) {
    this.test = test;
    this.or = function(other) {
      return new Predicate(t => this.test(t) || other.test(t));
    };
    this.and = function(other) {
      return new Predicate(t => this.test(t) && other.test(t));
    };
    this.negate = function() {
      return new Predicate(t => !this.test(t));
    };
  }

  function RealSet(start, end, rangeType, predF) {
    this.low = start;
    this.high = end;

    if (predF) {
      this.predicate = new Predicate(predF);
    } else {
      this.predicate = new Predicate(d => {
        switch (rangeType) {
          case RangeType.CLOSED:
            return start <= d && d <= end;
          case RangeType.BOTH_OPEN:
            return start < d && d < end;
          case RangeType.LEFT_OPEN:
            return start < d && d <= end;
          case RangeType.RIGHT_OPEN:
            return start <= d && d < end;
        }
      });
    }

    this.contains = function(d) {
      return this.predicate.test(d);
    };

    this.union = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.or(other.predicate).test(d)
      );
    };

    this.intersect = function(other) {
      var low2 = Math.min(this.low, other.low);
      var high2 = Math.max(this.high, other.high);
      return new RealSet(low2, high2, null, d =>
        this.predicate.and(other.predicate).test(d)
      );
    };

    this.subtract = function(other) {
      return new RealSet(this.low, this.high, null, d =>
        this.predicate.and(other.predicate.negate()).test(d)
      );
    };
  }
  set1 = new RealSet(set1.low, set1.high, set1.rangeType);
  set2 = new RealSet(set2.low, set2.high, set2.rangeType);
  var result = [];
  values.forEach(function(value) {
    result.push(set1[operation](set2).contains(value));
  });
  return result;
}
```
