---
id: 6995d0325fd7da40854fac40
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Em uma liĂ§ĂŁo anterior, vocĂȘ aprendeu sobre busca binĂĄria e como ela difere da busca linear.

Neste workshop, vocĂȘ colocarĂĄ isso em prĂĄtica construindo um algoritmo de busca binĂĄria.

Para comeĂ§ar, crie uma funĂ§ĂŁo `binarySearch` com os parĂąmetros `searchList` e `value`.

`searchList` Ă© o array de elementos dentro do qual a funĂ§ĂŁo vai buscar, e `value` Ă© o alvo que a funĂ§ĂŁo tentarĂĄ encontrar dentro do `searchList`.

# --hints--

VocĂȘ deve criar uma funĂ§ĂŁo `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Sua funĂ§ĂŁo `binarySearch` deve ter os parĂąmetros `searchList` e `value`.

```js
const regex = __helpers.functionRegex('binarySearch', ['searchList', 'value']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
