---
id: 6995d0325fd7da40854fac40
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Em uma lição anterior, você aprendeu sobre binary search e como ele difere do linear search.

Neste workshop, você colocará isso em prática construindo um algoritmo de busca binária.

Para começar, crie uma função `binarySearch` com os parâmetros `searchList` e `value`.

`searchList` é o array de elementos dentro do qual a função vai buscar, e `value` é o alvo que a função tentará encontrar dentro do `searchList`.

# --hints--

Você deve criar uma função `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Sua função `binarySearch` deve ter os parâmetros `searchList` e `value`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
