---
id: 6995d0325fd7da40854fac40
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Em uma lição anterior, você aprendeu sobre busca binária e como ela difere da busca linear.

Nesta oficina, você vai colocar isso em prática construindo um algoritmo de busca binária.

Para começar, crie uma função `binarySearch` com os parâmetros `searchList` e `value`.

`searchList` é o array de elementos dentro do qual a função vai buscar, e `value` é o alvo que a função tentará encontrar dentro do `searchList`.

# --hints--

Você deve criar uma função `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Sua função `binarySearch` deve ter os parâmetros `searchList` e `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
