---
id: 6996c7363a4a437839f1f15d
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora que você definiu os limites da sua busca, é hora de criar o loop que vai executar o algoritmo de busca binária.

A busca binária funciona reduzindo repetidamente o espaço de busca. Esse processo continua enquanto houver um intervalo válido de elementos para verificar.

Você pode expressar isso com um loop `while` que continuará enquanto o ponteiro `low` for menor ou igual ao ponteiro `high`.

Se `low` alguma vez se tornar maior que `high`, isso significa que o espaço de busca ficou vazio e o `value` não está no array.

Dentro da função `binarySearch`, crie um loop `while` com uma condição que verifica se `low` é menor ou igual a high.

# --hints--

Você deve criar um loop `while` que execute enquanto `low` for menor ou igual a `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
