---
id: 6996d1032feed88e2957dc8e
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Você pode ver que a última chamada imprime um array vazio. Para especificar melhor que o valor não está no array pesquisado, atualize a declaração `return` para retornar um array contendo dois elementos, um array vazio e a string `Value not found`.

# --hints--

Você deve atualizar a declaração `return` para retornar `[], 'Value not found'`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
    return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
