---
id: 6996c796385b6ea4e222d5e8
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Para começar com o loop, você precisa encontrar o índice do meio do espaço de busca atual.

Você pode obter o índice do meio calculando a média das variáveis `low` e `high`. Você terá que usar divisão inteira para obter a média após somar os dois valores, então a resposta será sempre um número inteiro arredondado para baixo, pois os índices devem ser inteiros.

Dentro do seu `while` loop, declare uma variável `mid` definida como a média de `low` e `high`.

# --hints--

Dentro do while loop você deve declarar uma variável `mid` definida como `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
