---
id: 6996ccadf9048c79b857bb15
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Você pode ver que a primeira chamada retorna `[3]`. Isso acontece porque `3` é o valor do meio em `[1, 2, 3, 4, 5]`.

Por outro lado, a segunda chamada retorna um array vazio. Isso ocorre porque, após verificar o valor do meio, o loop atualmente não tem a lógica para restringir o intervalo da busca atualizando as variáveis `low` ou `high`. Ele age apenas como uma verificação única do ponto médio e depois termina sem continuar a busca.

Para permitir que a busca binária continue restringindo sua busca, adicione uma declaração `else if` que verifique se `value` é maior que `valueAtMiddle`.

# --hints--

Você deve adicionar uma declaração `else if` com a condição `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
