---
id: 6996c95e5bede03bdf2937a2
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

A próxima coisa a fazer é criar uma condição que verificará se o `value` alvo está no meio.

Dentro do loop `while`, crie uma declaração `if` que verifique se o alvo `value` é igual a `valueAtMiddle`. Se for, retorne a variável `pathToTarget`.

# --hints--

Você deve criar uma declaração `if` com a condição `value === valueAtMiddle`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Você deve retornar a variável `pathToTarget` dentro da declaração `if`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
