---
id: 6996c675b599baa7b103192d
title: passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Na busca binária, existe o menor índice possível e o maior índice possível.

O menor índice possível representa o limite mais à esquerda na porção atual dos valores sendo pesquisados, neste caso, `searchList`. O maior índice possível, por outro lado, representa o limite mais à direita nos valores pesquisados.

Para considerar esses dois, você precisa levar em conta o array inteiro que está sendo pesquisado. Então, dentro da função `binarySearch`, defina uma variável `low` com o valor `0` e uma variável `high` com um valor que represente o último índice do array sendo pesquisado.

# --hints--

Você deve ter uma variável chamada `low` definida como `0` na sua função `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

Você deve ter uma variável chamada `high` definida como `searchList.length - 1` na sua função `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
