---
id: 6996c8f94fe6b250f11a2a1e
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora, inclua `valueAtMiddle` no array `pathToTarget` para acompanhar os passos dados durante a busca, independentemente de ser o alvo. Você pode fazer isso com o método `push()`.

# --hints--

Você deve usar o método `push()` para adicionar `valueAtMiddle` ao array `pathToTarget`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
