---
id: 6996c675b599baa7b103192d
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Na busca binária, existe o menor índice possível e o maior índice possível.

O menor índice possível representa o limite mais à esquerda na porção atual dos valores sendo pesquisados, neste caso, `searchList`. O maior índice possível, por outro lado, representa o limite mais à direita nos valores pesquisados.

Para considerar esses dois, você precisa levar em conta todo o array que está sendo pesquisado. Então, dentro da função `binarySearch`, declare uma variável `low` com o valor `0` e uma variável `high` com um valor que corresponda ao último índice do array pesquisado.

# --hints--

Você deve declarar uma variável chamada `low` com valor `0` na sua função `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Você deve declarar uma variável chamada `high` com valor `searchList.length - 1` na sua função `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
