---
id: 6996cfb6c01e9a802add38be
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Finalmente, se o `value` nĂŁo for encontrado no meio e tambĂ©m nĂŁo for maior que `valueAtMiddle`, entĂŁo o `value` deve ser menor que `valueAtMiddle` e deve estar Ă  esquerda. O bloco final `else` vai lidar com isso.

EntĂŁo, termine o loop removendo a palavra-chave `break` e adicionando um bloco `else`. Dentro dele, atualize a variĂĄvel `high` subtraindo `1` de `mid`.

# --hints--

VocĂȘ deve substituir a palavra-chave `break` por um bloco `else` que atualize a variĂĄvel `high` para `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
