---
id: 6996ca21a74104721f2c7e5d
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Você precisa testar as coisas para entender o fluxo do algoritmo nesta fase inicial.

Para isso, primeiro é preciso sair do loop. Isso porque a implementação atual permitirá apenas uma iteração, então se a condição não for satisfeita, haverá um loop infinito.

Logo após a declaração `if`, use a palavra-chave `break` para sair do loop `while`. Depois, após o loop `while`, retorne um array vazio para indicar que o valor não foi encontrado.

# --hints--

Você deve usar a palavra-chave `break` para sair do loop `while` após a declaração `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Você deve retornar um array vazio após seu loop `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
