---
id: 6996c60281feca6718cc15c1
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Dentro da sua função `binarySearch`, crie uma variável `pathToTarget` e inicialize-a como um array vazio, para começar.

# --hints--

Você deve ter uma variável `pathToTarget` definida como um array vazio `[]` dentro da sua função `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
