---
id: 6996c95e5bede03bdf2937a2
title: passo 8
challengeType: 1
dashedName: step-8
---

# --description--

O prĂłximo passo Ă© criar uma condiĂ§ĂŁo que verifique se o alvo `value` estĂĄ no meio.

Dentro do loop `while`, crie uma declaraĂ§ĂŁo `if` que verifique se o alvo `value` Ă© igual a `valueAtMiddle`. Se for, retorne a variĂĄvel `pathToTarget`.

# --hints--

VocĂȘ deve criar uma declaraĂ§ĂŁo `if` com a condiĂ§ĂŁo `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

VocĂȘ deve retornar a variĂĄvel `pathToTarget` dentro da declaraĂ§ĂŁo `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
