---
id: 6996cb5869acf58de751ae81
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Lembre-se de que a função recebe um array chamado `searchList` e o `value` que você está procurando.

Agora, faça a chamada da função com `binarySearch([1, 2, 3, 4, 5], 3)` e `binarySearch([1, 2, 3, 4, 5, 9], 4)` e registre ambas as declarações no console.

# --hints--

Você deve fazer a chamada da função com `binarySearch([1, 2, 3, 4, 5], 3)` e registrá-la no console.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

Você deve fazer a chamada da função com `binarySearch([1, 2, 3, 4, 5, 9], 4)` e registrá-la no console.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
