---
id: 6996d0a2c2ac55041145a13f
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Você percebeu que a segunda chamada da função agora tem os valores `3, 5, 4` na lista? Isso indica que a busca binária está funcionando como esperado.

Foi assim que aconteceu: o algoritmo primeiro verificou `3` como o meio da lista inicial. Como `4` é maior que `3`, a busca se deslocou para a metade direita. Em seguida, examinou `5` como o novo meio. Como `4` é menor que `5`, a busca se moveu para a esquerda, identificando finalmente `4` como o meio do intervalo final.

Para testar a função novamente, chame-a com `[1, 3, 5, 9, 14, 22], 10` e registre no console. Essa é uma situação em que o valor não será encontrado.

# --hints--

Você deve chamar a função com `binarySearch([1, 3, 5, 9, 14, 22], 10)` e registrar no console.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
