---
id: 6996c8375f4301004459ee85
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora que você calculou o índice `mid`, precisa realmente recuperar o valor do `searchList` nesse índice. Esse valor é o que você vai comparar com o parâmetro `value`, que é o alvo que você está buscando.

Crie uma variável `valueAtMiddle` e atribua a ela o elemento do `searchList` que está localizado no índice `mid`.

# --hints--

Você deve criar uma variável `valueAtMiddle` atribuída a `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
