---
id: 66ea8239b9eeb76f12299030
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Sehemu inayofuata ya warsha ni kukagua jinsi ya kuzalisha nambari ya bahati nasibu kati ya thamani mbili.

Anza kwa kuongeza `console.log()` nyingine inayorekodi ujumbe `"Now, generate a random number between two values."` kwenye konsoli.

Kisha, tengeneza kigezo kinachoitwa `min` na weka thamani ya `1` na tengeneza kigezo kinachoitwa `max` na weka thamani ya `100`.

Katika hatua inayofuata, utazalisha nambari ya bahati nasibu kati ya thamani hizi mbili.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Unapaswa kuwa na `console.log()` inayorekodi ujumbe `"Now, generate a random number between two values."`.

```js
assert.equal(getLogs()[3], "Now, generate a random number between two values.");
```

Unapaswa kuwa na kigezo kinachoitwa `min`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.min);
```

Kigezo chako cha `min` kinapaswa kuwa nambari.

```js
assert.isNumber(min);
```

Unapaswa kuweka nambari `1` kwenye kigezo chako cha `min`.

```js
assert.equal(min, 1);
```

Unapaswa kuwa na kigezo kinachoitwa `max`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.max);
```

Kigezo chako cha `max` kinapaswa kuwa nambari.

```js
assert.isNumber(max);
```

Unapaswa kuweka nambari `100` kwenye kigezo chako cha `max`.

```js
assert.equal(max, 100); 
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);
--fcc-editable-region--

--fcc-editable-region--
```
