---
id: 66ea845f54195c7020cc521b
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Usemi wa kihesabu wa kuzalisha nambari ya bahati nasibu kati ya thamani mbili ni ifuatayo:

```js
Math.random() * (maximum - minimum) + minimum;
```

Hii itatoa matokeo ambayo ni nambari ya desimali kati ya thamani mbili.

Tengeneza kigezo kinachoitwa `randomNum2` na ukaweke thamani yake kuwa matokeo ya kuzalisha thamani kati ya thamani za `min` na `max`.

Kisha, andika kigezo cha `randomNum2` kwenye konsoli ili kuona matokeo. Jaribu onyesha upya ukurasa kuona matokeo tofauti.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `randomNum2`.

```js
assert.isNotNull(randomNum2);
```

Kigezo chako cha `randomNum2` kinapaswa kuwa nambari.

```js
assert.isNumber(randomNum2);
```

Unapaswa kuweka thamani ya matokeo ya kuzalisha thamani kati ya `min` na `max`. Hakikisha hutoi thamani ya moja kwa moja (hardcode). Unapaswa kutumia usemi wa kihesabu uliotolewa hapo juu.

```js
assert.isAtLeast(randomNum2, min);
assert.isAtMost(randomNum2, max);

assert.notMatch(code, /randomNum2\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum2\s*=\s*\d+;?/);
```

Unapaswa kuandika kigezo cha `randomNum2` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*randomNum2\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;
--fcc-editable-region--

--fcc-editable-region--
```
