---
id: 66ea7fc048a6e26e33c598fe
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa, ni wakati wa kuzalisha nambari ya bahati nasibu kwa kutumia njia ya `Math.random()`.

Tengeneza kigezo kinachoitwa `randomNum` na weka thamani yake kuwa matokeo ya wito wa njia ya `Math.random()`.

Kisha, andika kigezo cha `randomNum` kwenye konsoli.

Jaribu kuongeza nafasi katika msimbo ili kuendesha tena bot na kuona nambari tofauti za bahati nasibu zinazozalishwa.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `randomNum`.

```js
assert.isNotNull(randomNum);
```

Unapaswa kuweka thamani ya matokeo ya wito wa njia ya `Math.random()` kwenye kigezo cha `randomNum`.

```js
assert.match(code, /randomNum\s*=\s*Math\.random\(\);?/);
```

Hupaswi kuweka nambari ya bahati nasibu moja kwa moja kwa kigezo cha `randomNum`. Hakikisha unatumia njia ya `Math.random()` kuzalisha nambari ya bahati nasibu.

```js
assert.notMatch(code, /randomNum\s*=\s*0\.\d+;?/);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomNum\s*=\s*\d+;?/);
```

Unapaswa kuandika kigezo cha `randomNum` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
