---
id: 66ea7fc048a6e26e33c598fe
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Sasa, ni wakati wa kuzalisha nambari ya bahati nasibu kwa kutumia njia ya `Math.random()`.

Tengeneza kigezo kinachoitwa `randomNum` na weka thamani yake kuwa matokeo ya wito wa njia ya `Math.random()`.

Kisha, andika kigezo cha `randomNum` kwenye konsoli.

Jaribu kuongeza nafasi katika msimbo ili kuendesha tena bot na kuona nambari tofauti za bahati nasibu zinazozalishwa.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.randomNum);
```

Unapaswa kuweka thamani ya matokeo ya wito wa njia ya `Math.random()` kwenye kigezo cha `randomNum`.

```js
const explorer = await __helpers.Explorer(code);
const randomNum = explorer.variables.randomNum;
assert.equal(randomNum?.value?.toString(), "Math.random()");
```

Unapaswa kuandika kigezo cha `randomNum` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*randomNum\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");
--fcc-editable-region--

--fcc-editable-region--
```
