---
id: 66eb488c0129e192542f0865
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na njia ya `Math.ceil()` kama ifuatavyo:

```js
const price = 10.01;
Math.ceil(price); // 11
```

Njia ya `Math.ceil()` inazungusha thamani hadi nambari kamili inayofuata karibu zaidi.

Tengeneza kigezo kinachoitwa `numRoundedUp` na weka thamani yake kuwa matokeo ya kuzungusha nambari ya desimali `3.2` hadi nambari kamili inayofuata karibu zaidi.

Kisha, andika kigezo cha `numRoundedUp` kwenye konsoli ili kuona matokeo.

# --hints--

Unapaswa kuwa na kigezo cha `numRoundedUp`.

```js
assert.isNotNull(numRoundedUp);
```

Kigezo cha `numRoundedUp` kinapaswa kuwa nambari.

```js
assert.isNumber(numRoundedUp);
```

Kigezo chako cha `numRoundedUp` kinapaswa kuwa sawa na matokeo ya kuzungusha nambari ya desimali `3.2` chini hadi nambari kamili inayokaribu zaidi.

```js
assert.equal(numRoundedUp, 4);
```

Hupaswi kuweka thamani ya `4` moja kwa moja kwa kigezo cha `numRoundedUp`.

```js
assert.notMatch(code, /numRoundedUp\s*=\s*4/);
```

Unapaswa kuandika kigezo cha `numRoundedUp` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*numRoundedUp\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

--fcc-editable-region--

--fcc-editable-region--
```
