---
id: 66eb3ec3b07dfe8f898646c0
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Kama ulivyojifunza katika mafundisho yaliyopita, njia ya `Math.floor()` huzungusha thamani chini hadi nambari kamili inayokaribu zaidi.

```js
const price = 10.99;
Math.floor(price); // 10
```

Tengeneza kigezo kinachoitwa `numRoundedDown` na weka thamani yake kuwa matokeo ya kuzungusha nambari ya desimali `6.7` chini hadi nambari kamili inayokaribu zaidi.

Kisha, andika kigezo cha `numRoundedDown` kwenye konsoli ili kuona matokeo.

# --hints--

Unapaswa kuwa na kigezo cha `numRoundedDown`.

```js
assert.isNotNull(numRoundedDown);
```

Kigezo cha `numRoundedDown` kinapaswa kuwa nambari.

```js
assert.isNumber(numRoundedDown);
```

Kigezo chako cha `numRoundedDown` kinapaswa kuwa sawa na matokeo ya kuzungusha nambari ya desimali `6.7` chini hadi nambari kamili inayokaribu zaidi.

```js
assert.equal(numRoundedDown, 6);
```

Hupaswi kuweka thamani ya `6` moja kwa moja kwa kigezo cha `numRoundedDown`.

```js
assert.notMatch(code, /numRoundedDown\s*=\s*6/);
```

Unapaswa kuandika kigezo cha `numRoundedDown` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*numRoundedDown\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");
--fcc-editable-region--

--fcc-editable-region--
```
