---
id: 66ece147323c91f654bf8dd9
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Katika mafundisho yaliyopita, ulijifunza jinsi ya kupata thamani kubwa zaidi na kiwango cha chini kutoka kwa safu ya nambari kama hii:

```js
Math.max(1, 2, 3, 4, 5); // 5
Math.min(1, 2, 3, 4, 5); // 1
```

Tengeneza kigezo kinachoitwa `maxNum` na ukae thamani yake kwa matokeo ya kupata nambari kubwa zaidi kati ya nambari `3, 125, 55, 24`. Kisha, andika thamani ya `maxNum` kwenye konsoli.

Chini ya hapo, tengeneza kigezo kinachoitwa `minNum` na ukae thamani yake kwa matokeo ya kupata nambari ya kiwango cha chini kati ya nambari `6, 90, 14, 90, 2`. Kisha, andika thamani ya `minNum` kwenye konsoli.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `maxNum`.

```js
assert.isNotNull(maxNum);
```

Kigezo chako cha `maxNum` kinapaswa kupewa thamani ya matokeo ya kupata nambari kubwa zaidi kati ya nambari `3, 125, 55, 24`.

```js
assert.equal(maxNum, 125);
```

Hupaswi kuweka thamani ya `125` moja kwa moja kwenye kigezo cha `maxNum`. Hakikisha unatumia njia ya `Math.max()` badala yake.

```js
assert.notMatch(code, /maxNum\s*=\s*125/);
```

Unapaswa kuandika thamani ya `maxNum` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*maxNum\s*\)/);
```

Unapaswa kuwa na kigezo kinachoitwa `minNum`.

```js
assert.isNotNull(minNum);
```

Kigezo chako cha `minNum` kinapaswa kupewa thamani ya matokeo ya kupata nambari ya kiwango cha chini kati ya nambari `6, 90, 14, 90, 2`.

```js
assert.equal(minNum, 2);
```

Hupaswi kuweka thamani ya `2` moja kwa moja kwenye kigezo cha `minNum`. Hakikisha unatumia njia ya `Math.min()` badala yake.

```js
assert.notMatch(code, /minNum\s*=\s*2/);
```

Unapaswa kuandika thamani ya `minNum` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*minNum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");

const randomInt = Math.floor(Math.random() * (max - min) + min);
console.log(randomInt);

console.log("The Math.ceil() method rounds the value up to the nearest whole integer.");

const numRoundedUp = Math.ceil(3.2);
console.log(numRoundedUp);

console.log(
  "The Math.round() method rounds the value to the nearest whole integer."
);

const numRounded = Math.round(2.7);
console.log(numRounded);
const numRounded2 = Math.round(11.2);
console.log(numRounded2);

console.log("The Math.max() and Math.min() methods are used to get the maximum and minimum number from a range.");

--fcc-editable-region--

--fcc-editable-region--
```
