---
id: 69218e46f4490b33f2481a1a
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Usemi wa kihesabu wa kuzalisha nambari kamili ya bahati nasibu kati ya thamani mbili ni ifuatayo:

```js
Math.floor(Math.random() * (maximum - minimum) + minimum);
```

Hii itatoa matokeo ambayo ni nambari kamili kati ya thamani mbili.

Tengeneza kigezo kinachoitwa `randomInt` na ukaweke thamani yake matokeo ya kuzalisha thamani kati ya thamani za `min` na `max`.

Kisha, andika kigezo cha `randomInt` ili kuona matokeo. Jaribu onyesha upya ukurasa kuona matokeo tofauti.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `randomInt`.

```js
assert.isNotNull(randomInt);
```

Kigezo chako cha `randomInt` kinapaswa kuwa nambari.

```js
assert.isNumber(randomInt);
```

Unapaswa kuweka thamani matokeo ya kuzalisha thamani kati ya `min` na `max`. Hakikisha hutoi thamani moja kwa moja kwenye msimbo. Unapaswa kutumia usemi wa kihesabu uliotolewa hapo juu.

```js
assert.isAtLeast(randomInt, min);
assert.isAtMost(randomInt, max);

// check for whole random numbers just in case a camper does this
assert.notMatch(code, /randomInt\s*=\s*\d+;?/);
```

Unapaswa andika kigezo cha `randomInt` kwenye konsoli.

```js
assert.match(code, /console\.log\(\s*randomInt\s*\);?/);
```


# --seed--

## --seed-contents--

```js
const botName = "MathBot";
const greeting = `Hi there! My name is ${botName} and I am here to teach you about the Math object!`;

console.log(greeting);

console.log("The Math.random() method returns a pseudo random number greater than or equal to 0 and less than 1.");

const randomNum = Math.random();
console.log(randomNum);

console.log("Now, generate a random number between two values.");

const min = 1;
const max = 100;

const randomNum2 = Math.random() * (max - min) + min;
console.log(randomNum2);

console.log("The Math.floor() method rounds the value down to the nearest whole integer.");

const numRoundedDown = Math.floor(6.7);
console.log(numRoundedDown);

console.log("Now, generate a random integer between two values.");
--fcc-editable-region--

--fcc-editable-region--
```
