---
id: 6842a6cd9836f0114a5b7a8a
title: Construir un Generador de Patrones Numéricos
challengeType: 27
dashedName: build-a-number-pattern-generator
---

# --description--

En este laboratorio practicarás los conceptos básicos de Python construyendo una pequeña aplicación que crea un patrón numérico.

**Objetivo:** Cumplir con los siguientes requisitos del usuario y lograr que pasen todas las pruebas para completar el laboratorio.

**Requisitos del usuario:**

1. Debes definir una función llamada `number_pattern` que tome un solo parámetro `n` (representando un entero positivo).
1. `number_pattern` debe usar un bucle `for`.
1. `number_pattern(n)` debe devolver una cadena con todos los enteros empezando desde 1 hasta `n` (incluido) separados por un espacio. Por ejemplo, `number_pattern(4)` debe devolver la cadena `1 2 3 4`.
1. Si el argumento pasado a la función no es un valor entero, la función debe devolver `Argument must be an integer value.`.
1. Si el argumento pasado a la función es menor que 1, la función debe devolver `Argument must be an integer greater than 0.`.

# --hints--

Debes tener una función `number_patter`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('number_pattern')
    `)
})
```

La función `number_pattern` debe tener un parámetro llamado `n`.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('number_pattern').has_args('n')
    `)
})
```

`number_pattern(4)` debe devolver `1 2 3 4`.

```js
({
    test: () => runPython(`
    assert number_pattern(4) == '1 2 3 4'
    `)
})
```

`number_pattern(12)` debe devolver `1 2 3 4 5 6 7 8 9 10 11 12`.

```js
({
    test: () => runPython(`
    assert number_pattern(12) == '1 2 3 4 5 6 7 8 9 10 11 12'
    `)
})
```

`number_pattern` debe devolver una lista de números separados por espacios para cualquier entero positivo.

```js
({
    test: () => runPython(`
    from random import randint
    n = randint(1,65)
    ls = ' '.join(map(lambda n: str(n), list(range(1, n+1))))
    assert number_pattern(n) == ls

    n2 = randint(34,122)
    ls2 = ' '.join(map(lambda n: str(n), list(range(1, n2+1))))
    assert number_pattern(n2) == ls2

    `)
})
```

`number_pattern` debe devolver `Argument must be an integer value.` cuando se le pasa un valor que no es un entero.

```js
({
    test: () => runPython(`
    string = 'Argument must be an integer value.'
    assert number_pattern('a') == string
    assert number_pattern(3.5) == string
    assert number_pattern(lambda n: n) == string
    `)
})
```

`number_pattern` debe devolver `Argument must be an integer greater than 0.` cuando se le pasa un entero no positivo.

```js
({
    test: () => runPython(`
    from random import randint
    string = 'Argument must be an integer greater than 0.'
    assert number_pattern(0) == string
    assert number_pattern(-1185) == string
    assert number_pattern(-45) == string
    assert number_pattern(randint(-9999999, 0)) == string
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def number_pattern(n):
    result = "1"
    if isinstance(n, int):
        if n > 0:
            for i in range(2, n + 1):
                result += f" {i}"
            return result
        return "Argument must be an integer greater than 0."
    return "Argument must be an integer value."

```
