---
id: 66a83601cd819e37f0dccd14
title: Paso 1
challengeType: 0
dashedName: step-1
---

# --description--

En este taller, practicarás trabajar con formularios HTML construyendo un Formulario de Retroalimentación del Hotel.

Para el texto introductorio, querrá mostrar el título principal seguido de una breve nota sobre la entrega de comentarios. 

Dentro de su elemento `body`, agregue un elemento `header`.

Dentro del elemento `header`, agregue un elemento `h1`. El texto del elemento `h1` debe ser `Hotel Feedback Form`.

Debajo de su elemento `h1`, agregue un elemento `p`. El texto del elemento <code>p` debe ser `Thank you for staying with us. Por favor, proporcione comentarios sobre su estadía reciente.</code>

# --hints--

Debe tener una etiqueta de apertura `header`.

```js
assert.match(code, /<header>/i);
```

Debe tener una etiqueta de cierre `header`.

```js
assert.match(code, /<\/header>/i);
```

Debes tener una etiqueta `h1` de apertura.

```js
assert.match(code, /<h1>/i);
```

Debes tener una etiqueta `h1` de cierre.

```js
assert.match(code, /<\/h1>/i);
```

Su elemento `h1` debe tener el texto `Hotel Feedback Form`.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.innerText.trim(), 'Hotel Feedback Form');
```

Su elemento `h1` debe estar dentro del `header` elemento.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.parentElement.tagName, 'HEADER');
```

Debe tener una etiqueta de apertura `p`.

```js
assert.match(code, /<p>/i);
```

Debe tener una etiqueta de cierre `p`.

```js
assert.match(code, /<\/p>/i);
```

Su elemento de párrafo debe tener el texto <code>Thank you for staying with us. Por favor, proporcione comentarios sobre su estadía reciente.</code>

```js
const pElement = document.querySelector('header p');
assert.strictEqual(pElement?.innerText.trim(), 'Thank you for staying with us. Please provide feedback on your recent stay.');
```

Tu elemento de párrafo debe estar dentro de tu `header`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.parentElement.tagName, 'HEADER');
```

Su elemento de párrafo debe venir después de su `h1`.

```js
const pElement = document.querySelector('p');
assert.strictEqual(pElement?.previousElementSibling.tagName, 'H1');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
