---
id: 66a9401c9d660d6bb15993e2
title: Paso 13
challengeType: 0
dashedName: step-13
---

# --description--

El formulario de comentarios de su hotel debe tener una opción para que los usuarios agreguen su edad. 

Comience agregando un elemento `label` con el texto `Age (optional):` al formulario.

El atributo `for` debe establecerse en `"age"`.

# --hints--

Deberías tener un elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 3);
```

Su elemento `label` debe tener un atributo `for` establecido en `"age"`.

```js
assert.isNotNull(document.querySelector('label[for="age"]'));
```

Su elemento `label` debe tener el texto `Age (optional):`.

```js
assert.strictEqual(document.querySelector('label[for="age"]')?.textContent.trim(), "Age (optional):");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
