---
id: 66a9577022877d72d8f43b4f
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

La siguiente sección del formulario preguntará a los usuarios por qué eligieron ese hotel en particular. Los usuarios tendrán la oportunidad de seleccionar múltiples opciones.

Comienza añadiendo otro elemento `fieldset`.

Dentro del elemento `fieldset`, agrega un elemento `legend` con el texto `Why did you choose to stay at our hotel? (Check all that apply)`.

# --hints--

Debería tener un elemento `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 3);
```

Deberías tener un elemento `legend` dentro del elemento `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(3) legend'), 1);
```

Tu elemento `legend` debe tener el texto de `Why did you choose to stay at our hotel? (Check all that apply)`. Verifica dos veces errores ortográficos y problemas de espaciado.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) legend')?.innerText.trim(), 'Why did you choose to stay at our hotel? (Check all that apply)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" value="yes" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" value="no" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
