---
id: 66a969951120be7818d8ee49
title: Pasos 24
challengeType: 0
dashedName: step-24
---

# --description--

La siguiente sección del formulario proporcionará a los usuarios la capacidad de dejar una calificación para el hotel. 

Comience añadiendo un nuevo elemento `fieldset` con un elemento `legend` anidado adentro. El `legend` debe tener el texto `Ratings`.

Debajo del elemento `legend`, agregue un elemento `label` con el texto `How was the service?`. El atributo `for` debe establecerse en `"service"`.

# --hints--

Debería tener un elemento `fieldset`.

```js
assert.lengthOf(document.querySelectorAll('fieldset'), 4);
```

Su `fieldset` debe tener un elemento `legend`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend'), 1);
```

Su legend debe tener el texto `Ratings`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend')[0]?.textContent.trim(), 'Ratings');
```

Su `fieldset` debe tener un elemento `label` debajo del `legend`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(4) legend + label'), 1);
```

Su `label` debe tener el texto `How was the service?`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.textContent.trim(), 'How was the service?');
```

Su `label` debe tener el atributo `for` establecido en `"service"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(4) legend + label')[0]?.getAttribute('for'), 'service');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>
          
          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>
          
          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        --fcc-editable-region--
        
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
