---
id: 66a93c95bc58e26a8fe95818
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Luego, agregue un `input` con el `type` `"email"` abajo de su `label` de correo electrónico. Este `input` debería tener un atributo `id` establecido al valor de `"email"`. Además, establezca el atributo `name` al valor de `"email"`. 

Este `input` también es obligatorio, así que asegúrese de añadir el atributo `required`.

Finalmente, añada un atributo `placeholder` configurado a `"example@email.com"`.

# --hints--

Debería tener un elemento `input` debajo de su `label` de correo electrónico.

```js
assert.isNotNull(document.querySelector('label[for="email"] + input'));
```

Su elemento `input` debería tener un `type` de `"email"`.

```js
assert.isNotNull(document.querySelector('input[type="email"]'));
```

Su elemento `input` debería tener un atributo `id` establecido en `"email"`.

```js
assert.isNotNull(document.querySelector('input#email'));
```

Su elemento `input` debería tener un atributo `name` establecido en `"email"`.

```js
assert.isNotNull(document.querySelector('input[name="email"]'));
```

Tu elemento `input` debe tener el atributo `required`.

```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('required'), "");
```

Su elemento `input` debería tener un atributo `placeholder` establecido en `"example@email.com"`.
  
```js
assert.strictEqual(document.querySelector('input#email')?.getAttribute('placeholder'), "example@email.com");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
        --fcc-editable-region--
          
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
