---
id: 66a954b2bcddba72076c1857
title: Paso 17
challengeType: 0
dashedName: step-17
---

# --description--

Debajo de tu elemento `label`, agrega un botón `radio` con el `id` establecido en `"no-option"` y el atributo name establecido en `"hotel-stay"`.

Debajo de tu nuevo botón `radio`, agrega otro elemento `label` con el atributo `for` establecido en `"no-option"`. El texto para el `label` debe ser `No`.

Cuando hayas terminado, ahora puedes probar los botones de radio seleccionando una opción a la vez.

# --hints--

Debes tener un elemento `input` debajo de tu primer elemento `label`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) label + input'));
```

Tu elemento `input` debe tener un atributo `type` con valor `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.getAttribute('type'), 'radio');
```

Tu elemento `input` debe tener un atributo `id` con el valor `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.id, 'no-option');
```

Tu elemento `input` debe tener un atributo `name` con el valor `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input')?.name, 'hotel-stay');
```

Debes tener un segundo elemento `label` debajo de tu segundo elemento `input`.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label'));
```

Tu elemento `label` debe tener un atributo `for` establecido en `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.htmlFor, "no-option");
```

Tu elemento `label` debe tener el texto `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input:nth-of-type(2) + label')?.textContent?.trim(), 'No');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          
          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
