---
id: 66a9824ac5d9f77ec304969f
title: Paso 34
challengeType: 0
dashedName: step-34
---

# --description--

Si desea que los usuarios tengan más espacio para escribir sus comentarios, puede usar un elemento `textarea`.

El elemento `textarea` es un control de entrada de texto de varias líneas que permite a los usuarios ingresar texto que es más largo que una sola línea. Se puede usar para crear una caja de comentarios, una entrada de mensaje u otras entradas de texto que requieran múltiples líneas.

Aquí hay un ejemplo de un elemento `textarea`:

```html
<textarea id="comments" name="comments" rows="4" cols="50"></textarea>
```

El atributo `rows` se usa para especificar la altura visible del `textarea`, y el atributo `cols` se usa para especificar el ancho visible del `textarea`.

Debajo de su elemento `label`, agregue un elemento `textarea`. En el próximo paso, agregarás los atributos necesarios.

# --hints--

Deberías tener un elemento `textarea` debajo de tu etiqueta.

```js
assert.exists(document.querySelector('label + textarea'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <input type="checkbox" id="ads" name="ads" value="ads" />
          <label for="ads">Social Media Ads</label>

          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          <label for="recommendation">Personal Recommendation</label>

          <input type="checkbox" id="location" name="location" value="location" />
          <label for="location">Location</label>

          <input
            checked
            type="checkbox"
            id="reputation"
            name="reputation"
            value="reputation"
          />
          <label for="reputation">Reputation</label>

          <input type="checkbox" id="price" name="price" value="price" />
          <label for="price">Price</label>
        </fieldset>

        <fieldset>
          <legend>Ratings</legend>

          <label for="service">How was the service?</label>

          <select name="service" id="service">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>

          <label for="food">How was the food?</label>

          <select name="food" id="food">
            <option value="poor">Poor</option>
            <option value="satisfactory">Satisfactory</option>
            <option value="good">Good</option>
            <option value="very-good">Very Good</option>
            <option selected value="excellent">Excellent</option>
          </select>
        </fieldset>

        --fcc-editable-region--
        <label for="comments">Other Comments?</label>

        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
