---
id: 66a9419e2d18476c645ce693
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

`input` de número se utiliza para crear un campo de entrada numérica. 

Aquí hay un ejemplo de un campo de entrada de número:

```html
<input type="number" id="age" name="age" min="18" max="100">
```

Los atributos `min` y `max` se usan para establecer los valores mínimo y máximo que se pueden ingresar en el campo de entrada.

Debajo de su elemento de `label`, añade un input con el atributo `type` establecido en `"number"` y un `id` de `"age"`.

El atributo `name` debe estar configurado como `"age"`, el atributo `min` debe ser `"3"` y el atributo `max` debe ser `"100"`.

# --hints--

Debe tener un `input` de número debajo de su elemento de `label`.

```js
assert.isNotNull(document.querySelector('label + input[type="number"]'));
```

Su `input` de número debe tener un `id` de `"age"`.

```js
assert.isNotNull(document.querySelector('label + input#age'));
```

Su `input` de número debe tener un atributo `name` establecido en `"age"`.

```js
assert.isNotNull(document.querySelector('label + input[name="age"]'));
```

Su `input` de número debe tener un atributo `min` establecido en `"3"`.

```js
assert.isNotNull(document.querySelector('label + input[min="3"]'));
```

Su `input` de número debe tener un atributo `max` establecido en `"100"`.

```js
assert.isNotNull(document.querySelector('label + input[max="100"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
      --fcc-editable-region--
          
      --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
