---
id: 66a93730719e1f68410cce54
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

El atributo `name` se usa para identificar los datos del formulario después de que se han enviado al servidor.

Aquí tienes un ejemplo de cómo usar el atributo `name`:

```html
<input type="email" name="email">
```

Agrega un atributo `name` al elemento `input` con el valor de `"name"`.

# --hints--

No debes eliminar el elemento `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

No debes crear un nuevo elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Tu elemento `input` debe tener un atributo `name` con el valor de `"name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('name'), 'name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name">
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
