---
id: 66a954b2bcddba72076c1857
title: Paso 20
challengeType: 0
dashedName: step-20
---

# --description--

Debajo de tu primer botón de radio, añade otro elemento `label` con el atributo `for` establecido en `"no-option"`. El texto de la etiqueta `label` debería ser `No`.

Debajo de tu segundo elemento `label`, añade un botón de tipo `radio` con el `id` establecido en `"no-option"`, y el atributo `name` establecido en `"hotel-stay"`.

Cuando hayas terminado, ahora puedes probar los botones de radio seleccionando una opción a la vez.

# --hints--

Deberías tener un elemento `label` debajo de tu primer botón de radio.

```js
assert.isNotNull(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label'));
```

Tu `label` debería tener un atributo `for` establecido en `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label')?.htmlFor, "no-option");
```

Tu `label` debería tener el texto `No`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) input[type="radio"] + label[for="no-option"]')?.textContent, 'No');
```

Deberías tener un segundo `input` debajo de tu segundo elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset:nth-of-type(2) label + input'), 2);
```

Tu `input` debería tener un `type` de `radio`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.getAttribute('type'), 'radio');
```

Tu `input` debería tener un `id` de `"no-option"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.id, 'no-option');
```

Tu `input` debería tener un `name` de `"hotel-stay"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(2) label + input:nth-of-type(2)')?.name, 'hotel-stay');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>
        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          --fcc-editable-region--
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
