---
id: 66a93bbe65a26169dbf3bc39
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

El formulario de retroalimentación de su hotel también debería recoger una dirección de correo electrónico del usuario. 

Comienza agregando un nuevo elemento `label` con el texto `Email address (required):` al formulario. Su elemento `label` debería tener un atributo `for` establecido en el valor de `"email"`.

# --hints--

Deberías tener un elemento `label`.

```js
assert.lengthOf(document.querySelectorAll('fieldset label'), 2);
```

Tu elemento `label` debería tener el texto `Email address (required):`.

```js
assert.strictEqual(document.querySelector('fieldset label:nth-of-type(2)')?.innerText.trim(), 'Email address (required):');
```

Su elemento `label` debería tener un atributo `for` establecido en el valor de `"email"`.

```js
assert.strictEqual(document.querySelector('fieldset label[for="email"]')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
--fcc-editable-region--
          
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
