---
id: 66a930b20f589b6664c51cb0
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Cuando un usuario proporciona su nombre completo, el `input` aceptará texto plano.

En lecciones anteriores, aprendiste cómo trabajar con el atributo `type` así:

```html
<input type="text">
```

Para tu elemento `input` existente, agrega un atributo `type` establecido en `"text"`. 

También dale al elemento `input` un atributo `id` con el valor de `"full-name"`.

# --hints--

No debes eliminar el elemento `input` existente.

```js
assert.isNotNull(document.querySelector('fieldset input'));
```

No debes crear un nuevo elemento `input`.

```js
assert.lengthOf(document.querySelectorAll('fieldset input'), 1);
```

Tu elemento `input` debe tener un atributo `type` con el valor de `"text"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('type'), 'text');
```

Tu elemento `input` debe tener un atributo `id` con el valor de `"full-name"`.

```js
const input = document.querySelector('fieldset input');
assert.strictEqual(input?.getAttribute('id'), 'full-name');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
        --fcc-editable-region--
          <input>
        --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
