---
id: 66a95d0eff8168747805f1f3
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Cuando deseas que un usuario seleccione múltiples opciones de una lista, puedes usar casillas de verificación. 

Aquí hay un ejemplo de cómo trabajar con casillas de verificación al tratar con opciones de alimentos:

```html
<fieldset>
  <legend>Food Options</legend>
  <input type="checkbox" id="pizza" name="food" value="pizza">
  <label for="pizza">Pizza</label>
  <input type="checkbox" id="burger" name="food" value="burger">
  <label for="burger">Burger</label>
</fieldset>
```

El atributo `value` se utiliza para especificar el valor que se enviará al servidor cuando el formulario se envíe.

Debajo de tu elemento `legend`, agrega un `input` de tipo checkbox con los atributos `id` y `value` establecidos en `"ads"`, y el atributo `name` establecido en `"choice"`.

Debajo de tu `input` de casilla de verificación, agrega un elemento `label` con el texto `Social Media Ads`. El atributo `for` debe establecerse en `"ads"`.

# --hints--

Debes tener un `input` de casilla de verificación debajo de tu elemento `legend`.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']"));
```

Tu `input` de casilla de verificación debe tener un atributo `id` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.id, "ads");
```

Tu `input` de casilla de verificación debe tener un atributo `name` establecido en `"choice"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.name, "choice");
```

Tu `input` de casilla de verificación debe tener un atributo `value` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) legend + input[type='checkbox']")?.value, "ads");
```

Debes tener un elemento `label` debajo de tu `input` de casilla de verificación.

```js
assert.exists(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label"));
```

Tu elemento `label` debe tener el texto `Social Media Ads`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.textContent?.trim(), "Social Media Ads");
```

Tu elemento `label` debe tener un atributo `for` establecido en `"ads"`.

```js
assert.strictEqual(document.querySelector("fieldset:nth-of-type(3) input[type='checkbox'] + label")?.getAttribute("for"), "ads");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="yes-option">Yes</label>
          <input id="no-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
        </fieldset>

        --fcc-editable-region--
        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>
          
        </fieldset>
        --fcc-editable-region--
      </form>
    </main>
  </body>
</html>
```
