---
id: 66a962954f4e0d76223b37ed
title: Pasos 24
challengeType: 0
dashedName: step-24
---

# --description--

A continuación, agrega otro elemento `label` con el texto de `Location` y el atributo `for` establecido en `"location"`.

Para el checkbox `input`, los atributos `id`, `name` y `value` deben establecerse en `"location"`.

Debajo de ese elemento `input`, añade otro elemento `label` con el texto de `Reputation` y el atributo `for` establecido en `"reputation "`.

Para el checkbox `input`, los atributos `id`, `name` y `value` deben establecerse en `"reputation"`.

# --hints--

Debes tener un elemento `label` con el texto de `Location`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.textContent, 'Location');
```

Tu `label` debe tener el atributo `for` establecido en `"location"`.

```js
assert.strictEqual(document.querySelectorAll('fieldset:nth-of-type(3) label')[2]?.getAttribute('for'), 'location');
```

Debes tener un tercer checkbox `input` después de tu `label`.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(3)[type='checkbox']"));
```

Tu tercer checkbox `input` debe tener un `id` establecido en `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('id'), 'location');
```

Tu tercer checkbox debe tener un `name` de `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('name'), 'location');
```

Tu tercer checkbox debe tener un `value` de `"location"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[2]?.getAttribute('value'), 'location');
```

Debes tener un elemento `label` con el texto de `Reputation` debajo del checkbox location.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.textContent?.trim(), 'Reputation');
```

Tu `label` debe tener el atributo `for` establecido en `"reputation"`.

```js
assert.strictEqual(document.querySelector('fieldset:nth-of-type(3) input[value="location"] + label')?.getAttribute('for'), 'reputation');
```

Debes tener un cuarto checkbox `input` debajo de tu `label`.

```js 
assert.exists(document.querySelector("fieldset:nth-of-type(3) label + input:nth-of-type(4)[type='checkbox']"));
```

Debes tener un cuarto checkbox `input` con un `id` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('id'), 'reputation');
```

Tu cuarto checkbox debe tener un `name` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('name'), 'reputation');
```

Tu cuarto checkbox debe tener un `value` de `"reputation"`.

```js
assert.strictEqual(document.querySelectorAll("fieldset:nth-of-type(3) label + input[type='checkbox']")[3]?.getAttribute('value'), 'reputation');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
          <input type="text" id="full-name" name="name" placeholder="Ex. John Doe" required size="20">

          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
            size="20"
          />
          <label for="age">Age (optional):</label>
          <input type="number" name="age" id="age" min="3" max="100" />
        </fieldset>

        <fieldset>
          <legend>Was this your first time at our hotel?</legend>
          <label for="yes-option">Yes</label>
          <input id="yes-option" type="radio" name="hotel-stay" />
          <label for="no-option">No</label>
          <input id="no-option" type="radio" name="hotel-stay" />
        </fieldset>

        <fieldset>
          <legend>
            Why did you choose to stay at our hotel? (Check all that apply)
          </legend>

          <label for="ads">Social Media Ads</label>
          <input type="checkbox" id="ads" name="ads" value="ads" />

          <label for="recommendation">Personal Recommendation</label>
          <input
            type="checkbox"
            id="recommendation"
            name="recommendation"
            value="recommendation"
          />
          --fcc-editable-region--

          --fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
