---
id: 67a51d3a8a6fe123b77b0c6e
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Los elementos `input` pueden tener un atributo `size`. Este atributo define el número de caracteres que deberían ser visibles mientras el usuario escribe en el input. El valor de `size` debe ser un entero no negativo mayor que cero. Si no se especifica `size`, o se especifica con un valor no válido, el input tendrá el ancho predeterminado establecido por el navegador.

```html
<label for="lastName">Last Name:</label>
<input id="lastName" name="lastName" type="text" size="10" />
```

Asigne a los inputs de nombre y correo electrónico un atributo `size` con un valor de `"20"`.

# --hints--

Debería asignar al `input` de nombre un atributo `size` con un valor de `"20"`.

```js
const label1 = document.querySelector('label:nth-of-type(1) + input');
assert.strictEqual(label1?.getAttribute('size'), '20');
```

Debería asignar al `input` de correo electrónico un atributo `size` con un valor de `"20"`.

```js
const label2 = document.querySelector('label:nth-of-type(2) + input');
assert.strictEqual(label2?.getAttribute('size'), '20');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Hotel Feedback Form</title>
  </head>
  <body>
    <header>
      <h1>Hotel Feedback Form</h1>
      <p>
        Thank you for staying with us. Please provide feedback on your recent
        stay.
      </p>
    </header>
    <main>
      <form method="POST" action="https://hotel-feedback.freecodecamp.org">
        <fieldset>
          <legend>Personal Information</legend>
          <label for="full-name">Name (required):</label>
--fcc-editable-region--
          <input type="text" id="full-name" name="name" placeholder="e.g., John Doe" required>
          <label for="email">Email address (required):</label>
          <input
            placeholder="example@email.com"
            required
            id="email"
            type="email"
            name="email"
          />         
--fcc-editable-region--
        </fieldset>
      </form>
    </main>
  </body>
</html>
```
