---
id: 6909ed80e1c195f49564edac
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop lernen Sie, wie Sie mit bedingten Anweisungen und Vergleichsoperatoren arbeiten.

Beginnen Sie damit, eine Variable namens `truthyOrFalsy` zu deklarieren und ihr den Boolean-Wert `true` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `truthyOrFalsy` haben.

```js
assert.exists(truthyOrFalsy);
```

Sie sollten `true` Ihrer Variable `truthyOrFalsy` zuweisen.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
