---
id: 690a06833d3e83b5226ade78
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Eine bedingte Anweisung wird verwendet, um Code abhängig von einer Bedingung auszuführen. Sie besteht aus dem `if`-Schlüsselwort, gefolgt von einer in runde Klammern eingeschlossenen Bedingung und einem in geschweifte Klammern eingeschlossenen Anweisungsblock. Hier ist ein Beispiel für eine `if`-Anweisung:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Wenn `condition` ein truthy-Wert ist (nicht-leerer String, von Null verschiedener Wert, `true` usw.), wird der Code im Anweisungsblock der `if`-Anweisung ausgeführt.

Erstellen Sie unter Ihrer Variablen eine `if`-Anweisung, die `hasDeveloperJob` als Bedingung verwendet. Protokollieren Sie innerhalb des Anweisungsblocks Ihrer `if`-Anweisung `"Timmy is employed as a developer."` in der Konsole.

# --hints--

Sie sollten eine `if`-Anweisung haben.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Ihre `if`-Anweisung sollte `hasDeveloperJob` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Sie sollten `"Timmy is employed as a developer."` innerhalb Ihrer `if`-Anweisung in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
