---
id: 690a57fef66aac39dffc61a2
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Da Sie nun wissen, was truthy- und falsy-Werte sind, entfernen Sie sowohl Ihre Variablendeklaration als auch die `console.log`-Anweisung aus Ihrem Code.

# --hints--

Sie sollten keine `truthyOrFalsy`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

Sie sollten `console.log(Boolean(truthyOrFalsy))` nicht in Ihrem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
