---
id: 690a110a76d0496d9e7c28eb
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Deklarieren Sie eine weitere Variable namens `timmyAge` und weisen Sie ihr den Wert `18` zu.

# --hints--

Sie sollten eine Variable namens `timmyAge` haben.

```js
assert.exists(timmyAge);
```

Sie sollten `18` Ihrer Variable `timmyAge` zuweisen.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
