---
id: 690a3f884e083a6a33b81fcf
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Ändern Sie schließlich die Deklaration von `timmyAge` und setzen Sie sie auf eine Zahl, die kleiner als `16` ist. Danach sehen Sie `"Timmy is not old enough to drive."` in der Konsole ausgegeben.

# --hints--

Ihre `timmyAge`-Variable sollte einen Wert haben, der kleiner als `16` ist.

```js
assert.isAtMost(timmyAge, 15);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--
const timmyAge = 18;
--fcc-editable-region--

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```

# --solutions--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 15;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
} else {
  console.log("Timmy is not old enough to drive.");
}
```
