---
id: 690a562fd611e8155599ce98
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Ändern Sie die Zuweisung von `truthyOrFalsy`, indem Sie ihren Wert auf `"freeCodeCamp"` setzen. Sie werden weiterhin `true` in der Konsole sehen, da ein nicht-leerer String ein truthy-Wert ist.

# --hints--

Ihre `truthyOrFalsy`-Variable sollte den Wert `"freeCodeCamp"` haben.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
