---
id: 690a57095c4afb27bbe751b3
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Falsy-Werte sind Werte, die in einem Boolean-Kontext zu `false` ausgewertet werden. Beispiele für falsy-Werte sind:

- `""` (leerer String)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Sie werden im Verlauf dieses Kurses alle kennenlernen.

Ändern Sie vorerst die Zuweisung von `truthyOrFalsy`, indem Sie seinen Wert auf einen leeren String setzen. Danach wird `false` in der Konsole ausgegeben.

# --hints--

Sie sollten `truthyOrFalsy` einen leeren String zuweisen.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
