---
id: 690a0f7a5c22c15222a6aa5c
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Erstellen Sie nun eine weitere `if`-Anweisung, die `isTimmyAGamer` als Bedingung verwendet. Protokollieren Sie innerhalb des Körpers Ihrer `if`-Anweisung `"Timmy loves to play World of Warcraft."` in der Konsole.

Beachten Sie, dass in der Konsole nichts protokolliert wird, da die Bedingung Ihrer neuen `if`-Anweisung falsch ist.

# --hints--

Sie sollten zwei `if`-Anweisungen in Ihrem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Ihre zweite `if`-Anweisung sollte `isTimmyAGamer` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Sie sollten `"Timmy loves to play World of Warcraft."` innerhalb Ihrer zweiten `if`-Anweisung in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
