---
id: 690a3ad44874a928a2fe3077
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Vergleichsoperatoren ermöglichen es Ihnen, Werte zu vergleichen und komplexere bedingte Anweisungen zu schreiben.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Zum Beispiel gibt der obige Code `"Happy Birthday!"` nur dann in der Konsole aus, wenn die Werte der Variablen `today` und `birthday` gleich sind.

Erstellen Sie eine dritte `if`-Anweisung. Verwenden Sie für deren Bedingung den `>=`-Operator, um zu prüfen, ob `timmyAge` größer oder gleich `16` ist.

Geben Sie im Rumpf Ihrer neuen `if`-Anweisung `"Timmy is old enough to drive."` in der Konsole aus.

# --hints--

Sie sollten eine dritte `if`-Anweisung in Ihrem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Ihre dritte `if`-Anweisung sollte `timmyAge >= 16` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

Sie sollten `"Timmy is old enough to drive."` in der Konsole innerhalb Ihrer dritten `if`-Anweisung ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
