---
id: 690a58cc4f828a49eab1b53a
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Deklarieren Sie eine Variable mit dem Namen `hasDeveloperJob` und weisen Sie ihr den Boolean-Wert `true` zu.

# --hints--

Sie sollten eine Variable mit dem Namen `hasDeveloperJob` haben.

```js
assert.exists(hasDeveloperJob);
```

Sie sollten `true` Ihrer Variable `hasDeveloperJob` zuweisen.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
