---
id: 690a0df27ee9ee36d935d530
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Deklarieren Sie eine weitere Variable namens `isTimmyAGamer` und weisen Sie ihr den Boolean `false` zu.

# --hints--

Sie sollten eine Variable namens `isTimmyAGamer` deklarieren.

```js
assert.exists(isTimmyAGamer);
```

Sie sollten der Variable `false` den Boolean-Wert `isTimmyAGamer` zuweisen.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
