---
id: 690a3ce03013be44d116136e
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Eine bedingte Anweisung kann eine `else`-Klausel haben, die Code ausführt, wenn die `if`-Bedingung falsch ist. Hier ist ein Beispiel für eine `if...else`-Anweisung:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

Fügen Sie der bestehenden `else`-Anweisung eine `if`-Klausel hinzu. Protokollieren Sie innerhalb des Körpers Ihrer `else`-Klausel `"Timmy is not old enough to drive."` in der Konsole.

# --hints--

Ihre dritte `if`-Anweisung sollte eine `else`-Klausel haben.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{/s);
```

Sie sollten `"Timmy is not old enough to drive."` innerhalb des Körpers Ihrer `else`-Klausel in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s+else\s+{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\2\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
