---
id: 690a51806fb743bc896c3af7
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Wie Sie in einer vorherigen Lektion gelernt haben, sind Booleans ein Datentyp mit nur `true` und `false` Werten.

Ein truthy-Wert ist ein Wert, der in einem Boolean-Kontext als `true` ausgewertet wird. Beispiele für truthy-Werte sind:

- Nicht-leere Strings
- Zahlen ungleich null
- Der Boolean `true`

Und andere Werte, denen Sie später begegnen werden.

Sie können die Funktion `Boolean()` verwenden, um die truthy-Eigenschaft eines Werts zu überprüfen. Zum Beispiel gibt `Boolean("Hello World!")` `true` zurück, weil `"Hello World!"` truthy ist.

Rufen Sie nun `Boolean(truthyOrFalsy)` auf und geben Sie das Ergebnis in der Konsole aus.

# --hints--

Sie sollten `Boolean(truthyOrFalsy)` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
