---
id: 61329d52e5010e08d9b9d66b
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

접근성과 SEO에 중요한 또 다른 `meta` 요소는 `description` 정의입니다. `content` 속성의 값은 검색 엔진이 페이지 설명을 제공하는 데 사용됩니다.

`meta` 속성을 `name`로 설정한 `description` 요소를 추가하고 유용한 `content` 속성을 지정하세요.

# --hints--

새로운 `meta` 요소를 `head`에 추가하세요.

```js
assert.equal(document.querySelectorAll('meta').length, 3);
```

`meta`에 `name` 값을 가진 `description` 속성을 지정하세요.

```js
assert.exists(document.querySelector('meta[name="description"]'));
```

`meta`에 `content` 속성을 지정하세요.

```js
assert.notEmpty(document.querySelector('meta[name="description"]')?.content);
```

`content` 속성 값은 165자를 넘지 않아야 합니다. _이것이 구글의 최대 설명 길이입니다._

```js
assert.isAtMost(document.querySelector('meta[name="description"]')?.content?.length, 165);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
