---
id: 6133acc353338c0bba9cb553
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

마지막으로 `head`에서 `title` 요소는 화면 읽기 프로그램이 페이지 내용을 이해하는 데 유용합니다. 또한, _SEO_의 중요한 부분입니다.

페이지에 설명적이고 간결한 `title`를 지정하세요.

# --hints--

`title`에 `head` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('head > title'));
```

`title`는 60자를 넘지 않도록 하세요.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

`title` 요소를 좀 더 설명적으로 작성해 보세요. _힌트: 최소 15자 이상_

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
