---
id: 6133d11ef548f51f876149e3
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

내비게이션은 접근성의 핵심이며, 화면 읽기 프로그램은 페이지 구조를 제공하는 데 의존합니다. 이는 의미론적 HTML 요소로 구현됩니다.

페이지에 `header`와 `main` 요소를 추가하세요.

`header` 요소는 페이지를 소개하고 내비게이션 메뉴를 제공하는 데 사용됩니다.

`main` 요소는 페이지의 핵심 콘텐츠를 포함합니다.

# --hints--

`header`에 `body` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('body > header'));
```

`main`에 `body` 요소를 추가해야 합니다.

```js
assert.exists(document.querySelector('body > main'));
```

`header` 요소는 `main` 요소보다 앞에 와야 합니다.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    
  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
