---
id: 6140883f74224508174794c0
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`header`가 부모 컨테이너의 전체 너비를 차지하도록 만들고, `height`을 `50px`로 설정하세요. 그리고 `background-color`을 `#1b1b32`로 설정하세요. 그런 다음 `display`가 _Flexbox_를 사용하도록 설정하세요.

# --hints--

`header` 요소 선택자를 사용하세요.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('header'));
```

`header`에 `width`를 `100%`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.width, '100%');
```

`header`에 `height`를 `50px`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.height, '50px');
```

`header`에 `background-color`를 `#1b1b32`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.backgroundColor, 'rgb(27, 27, 50)');
```

`header`에 `display`를 `flex`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('header')?.display, 'flex');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

```
