---
id: 6140827cff96e906bd38fc2b
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

<a href="https://design-style-guide.freecodecamp.org/" target="_blank" rel="noopener noreferrer nofollow">freeCodeCamp 스타일 가이드</a>에 설명된 대로, 로고는 `35 / 4`의 종횡비를 유지하고 텍스트 주변에 패딩이 있어야 합니다.

먼저, 로고가 보이도록 `background-color`를 `#0a0a23`으로 변경하세요. 그런 다음, 원하는 종횡비를 `35 / 4`로 설정하기 위해 `aspect-ratio` 속성을 사용하세요. 마지막으로, 모든 면에 `0.4rem`의 `padding`을 추가하세요.

# --hints--

`#logo`에 `background-color`를 `#0a0a23`으로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.backgroundColor, 'rgb(10, 10, 35)');
```

`aspect-ratio` 속성을 사용하세요.

```js
assert.notEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio);
```

`height` 속성은 사용하지 마세요.

```js
assert.isEmpty(new __helpers.CSSHelp(document).getStyle('#logo')?.height);
```

`width` 속성은 변경하지 마세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

`img`에 `35 / 4`의 `aspect-ratio`를 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.aspectRatio, '35 / 4');
```

`img`에 `padding`를 `0.4rem`만큼 주어야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.padding, '0.4rem');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

#logo {
  width: max(10rem, 18vw);
  --fcc-editable-region--
  
  --fcc-editable-region--
}

```
