---
id: 614202874ca576084fca625f
title: Step 16
challengeType: 0
dashedName: step-16
---

# --description--

모든 `region` 역할에는 화면 읽기 프로그램 사용자가 영역의 목적을 이해할 수 있도록 돕는 레이블이 필요합니다. 레이블을 추가하는 한 가지 방법은 영역 안에 제목 요소를 넣고 `aria-labelledby` 속성으로 참조하는 것입니다.

다음 `aria-labelledby` 속성들을 `section` 요소에 추가하세요:

- `student-info`
- `html-questions`
- `css-questions`

그런 다음 각 `section` 요소 안에, 해당 `h2` 속성과 일치하는 `id`를 가진 `aria-labelledby` 요소를 중첩하세요. 각 `h2`에 적절한 텍스트 내용을 넣으세요.

# --hints--

첫 번째 `section` 요소에는 `aria-labelledby` 속성으로 `student-info`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

두 번째 `section` 요소에는 `aria-labelledby` 속성으로 `html-questions`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

세 번째 `section` 요소에는 `aria-labelledby` 속성으로 `css-questions`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

첫 번째 `h2` 요소 안에 `section` 요소를 중첩해야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

두 번째 `h2` 요소 안에 `section` 요소를 중첩해야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

세 번째 `h2` 요소 안에 `section` 요소를 중첩해야 합니다.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

첫 번째 `h2` 요소에는 `id` 속성으로 `student-info`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

두 번째 `h2` 요소에는 `id` 속성으로 `html-questions`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

세 번째 `h2` 요소에는 `id` 속성으로 `css-questions`를 주어야 합니다.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

첫 번째 `h2` 요소에 적절한 텍스트 내용을 넣으세요. _힌트: `Student Info`를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

두 번째 `h2` 요소에 적절한 텍스트 내용을 넣으세요. _힌트: `HTML`를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

세 번째 `h2` 요소에 적절한 텍스트 내용을 넣으세요. _힌트: `CSS`를 선택했을 것입니다_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
