---
id: 6143908b6aafb00a659ca189
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

폼 입력에 대한 모범 사례를 염두에 두고, 각 `input`에 적절한 `type`와 `name` 속성을 지정하세요. 그런 다음, 첫 번째 `input`에 `placeholder` 속성을 지정하세요.

# --hints--

첫 번째 `input`에 `type`가 `text`인 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.type, 'text');
```

두 번째 `input`에 `type`가 `email`인 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.type, 'email');
```

세 번째 `input`에 `type`가 `date`인 속성을 지정해야 합니다.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.type, 'date');
```

첫 번째 `input`에 적절한 `name` 속성을 지정해야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('input')?.[0]?.name?.length, 1);
```

두 번째 `input`에 적절한 `name` 속성을 지정해야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('input')?.[1]?.name?.length, 1);
```

세 번째 `input`에 적절한 `name` 속성을 지정해야 합니다.

```js
assert.isAtLeast(document.querySelectorAll('input')?.[2]?.name?.length, 1);
```

첫 번째 `input`에 `placeholder` 속성을 지정해야 합니다.

```js
assert.notEmpty(document.querySelectorAll('input')?.[0]?.placeholder);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label for="student-name">Name:</label>
            <input id="student-name" />
          </div>
          <div class="info">
            <label for="student-email">Email:</label>
            <input id="student-email" />
          </div>
          <div class="info">
            <label for="birth-date">Date of Birth:</label>
            <input id="birth-date" />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
