---
id: 6143639d5eddc7094161648c
title: Step 20
challengeType: 0
dashedName: step-20
---

# --description--

각 `input`를 해당 `label` 요소와 연결하는 것이 중요합니다. 이렇게 하면 보조 기술 사용자가 입력에 대한 시각적 참조를 얻을 수 있습니다.

이는 `label`에 `for` 속성을 부여하여 수행하며, 이 속성에는 `id`의 `input`가 포함됩니다.

이 섹션에서는 학생의 이름, 이메일 주소, 생년월일을 입력받습니다. `label` 요소에 적절한 `for` 속성과 텍스트 내용을 부여하세요. 그런 다음 `input` 요소를 해당 `label` 요소와 연결하세요.

# --hints--

첫 번째 `label` 요소에 적절한 `for` 속성을 부여하세요.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[0]?.htmlFor?.length, 1);
```

두 번째 `label` 요소에 적절한 `for` 속성을 부여하세요.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[1]?.htmlFor?.length, 1);
```

세 번째 `label` 요소에 적절한 `for` 속성을 부여하세요.

```js
assert.isAtLeast(document.querySelectorAll('label')?.[2]?.htmlFor?.length, 1);
```

첫 번째 `label` 요소에 적절한 텍스트 내용을 부여하세요.

```js
assert.isNotEmpty(document.querySelectorAll('label')[0]?.textContent.trim());
```

두 번째 `label` 요소에 적절한 텍스트 내용을 부여하세요.

```js
assert.isNotEmpty(document.querySelectorAll('label')[1]?.textContent.trim());
```

세 번째 `label` 요소에 적절한 텍스트 내용을 부여하세요.

```js
assert.isNotEmpty(document.querySelectorAll('label')[2]?.textContent.trim());
```

첫 번째 `input` 요소에 첫 번째 `id`의 `for` 속성과 일치하는 `label` 속성을 부여하세요.

```js
assert.equal(document.querySelectorAll('input')?.[0]?.id, document.querySelectorAll('label')?.[0]?.htmlFor);
```

두 번째 `input` 요소에 두 번째 `id`의 `for` 속성과 일치하는 `label` 속성을 부여하세요.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.id, document.querySelectorAll('label')?.[1]?.htmlFor);
```

세 번째 `input` 요소에 세 번째 `id`의 `for` 속성과 일치하는 `label` 속성을 부여하세요.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.id, document.querySelectorAll('label')?.[2]?.htmlFor);
```

하나의 `id` 요소에 대해 같은 `input` 속성을 두 번 이상 사용하지 마세요.

```js
const id = (n) => document.querySelectorAll('input')?.[n]?.id;
assert.notEqual(id(0), id(1));
assert.notEqual(id(0), id(2));
assert.notEqual(id(1), id(2));
```

하나의 `for` 요소에 대해 같은 `label` 속성을 두 번 이상 사용하지 마세요.

```js
const htmlFor = (n) => document.querySelectorAll('label')?.[n]?.htmlFor;
assert.notEqual(htmlFor(0), htmlFor(1));
assert.notEqual(htmlFor(0), htmlFor(2));
assert.notEqual(htmlFor(1), htmlFor(2));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a href="#student-info">INFO</a></li>
          <li><a href="#html-questions">HTML</a></li>
          <li><a href="#css-questions">CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
--fcc-editable-region--
        <section role="region" aria-labelledby="student-info">
          <h2 id="student-info">Student Info</h2>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
          <div class="info">
            <label></label>
            <input />
          </div>
        </section>
--fcc-editable-region--
        <section role="region" aria-labelledby="html-questions">
          <h2 id="html-questions">HTML</h2>
        </section>
        <section role="region" aria-labelledby="css-questions">
          <h2 id="css-questions">CSS</h2>
        </section>
      </form>
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}

h1,
h2 {
  font-family: Verdana, Tahoma;
}

h2 {
  border-bottom: 4px solid #dfdfe2;
}

```
