---
id: 614ccc21ea91ef1736b9b578
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 기본 보일러플레이트가 제공되었습니다.

접근성 여정을 시작하려면 `lang` 요소에 `html` 속성을 추가하세요. 이렇게 하면 화면 읽기 프로그램이 페이지의 언어를 식별하는 데 도움이 됩니다.

또한 스타일시트를 페이지에 연결하는 것도 잊지 마세요.

# --hints--

`html` 요소에 `lang` 속성을 부여해야 합니다. _힌트: 영어의 경우 `en` 값을 사용할 수 있습니다._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

`link` 요소 안에 `head` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("head>link"));
```

`link` 요소에는 `rel="stylesheet"`와 `href="styles.css"`가 있어야 합니다.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
