---
id: 66325a250690a3612c1db0f6
title: Passo 5
challengeType: 0
dashedName: step-5
---

# --description--

CamperBot criou uma nova variável chamada `bgHexCodeSpanElement` para armazenar a referência ao elemento `span` com o `id` `bg-hex-code`.No entanto, quando eles tentam registrar essa variável no console, eles obtêm `null`. 

`null` é um valor especial em JavaScript que representa a ausência de um valor.Isso pode acontecer quando você tenta acessar uma propriedade de um objeto que não existe. 

Neste caso, o CamperBot não está passando o seletor correto para o método `document.querySelector`.

Corrija a linha `document.querySelector("bg-hex-code")` para que ela selecione corretamente o elemento com o `id` `bg-hex-code`.

# --hints--

Você deve corrigir `"bg-hex-code"` para que ele referencie corretamente o `id` de `bg-hex-code`.Lembre-se de que os seletores de `id` em CSS começam com um símbolo de hashtag (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

--fcc-editable-region--
const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("bg-hex-code");

console.log(bgHexCodeSpanElement);
--fcc-editable-region--
```
