---
id: 6650c9a94d6e13d14a043a69
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CamperBot está tentando criar um trocador aleatório de cor de fundo.Mas eles continuam enfrentando problemas e precisam da sua ajuda para depurar o código. 

CamperBot já adicionou o HTML e o CSS para o projeto.Mas eles estão confusos sobre por que nenhum dos estilos e conteúdos está aparecendo na página.

Quando eles abrem o console, veem esta mensagem:

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

Erros de sintaxe são lançados quando o motor JavaScript encontra algo que não consegue interpretar.Neste caso, parece que o CamperBot tem erros de sintaxe no array `darkColorsArr`.

Corrija os erros de sintaxe no array `darkColorsArr` e você deverá ver o conteúdo e os estilos aparecerem na página.

# --hints--

Você deve resolver os erros de sintaxe no array `darkColorsArr`.Lembre-se de que os elementos do array devem ser separados por vírgulas.

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
