---
id: 66ed41f912d0bb1dc62da5dd
title: بناء بوت أسئلة وأجوبة بلغة JavaScript
challengeType: 26
dashedName: lab-javascript-trivia-bot
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تسجل في وحدة التحكم الرسالة `"Hello! I'm your coding fun fact guide!"` كتحية للمستخدم.
1. يجب أن تنشئ ثلاثة متغيرات: `botName` و `botLocation` و `favoriteLanguage`، تخزن اسم البوت، ومكانه، ولغة البرمجة المفضلة لديه على التوالي.
1. يجب أن تسجل في وحدة التحكم الرسالة `"My name is (botName) and I live on (botLocation)."`.
1. يجب أن تسجل في وحدة التحكم الرسالة `"My favorite programming language is (favoriteLanguage)."`.
1. يجب أن تستخدم `let` لإنشاء متغير `codingFact` وتعيين سلسلة نصية له تحتوي على حقيقة ممتعة عن لغة البرمجة المفضلة للبوت، مع تضمين استخدام المتغير `favoriteLanguage`.
1. يجب أن تسجل المتغير `codingFact` في وحدة التحكم.
1. يجب أن تعيد تعيين المتغير `codingFact` إلى حقيقة جديدة عن لغة البوت المفضلة باستخدام المتغير `favoriteLanguage` مرة أخرى.
1. يجب أن تسجل المتغير `codingFact` في وحدة التحكم مرة أخرى.
1. يجب أن تعيد تعيين المتغير `codingFact` مرة أخرى إلى حقيقة جديدة أخرى عن لغة البوت المفضلة باستخدام المتغير `favoriteLanguage`.
1. يجب أن تسجل المتغير `codingFact` في وحدة التحكم للمرة الثالثة.
1. يجب أن تسجل في وحدة التحكم الرسالة `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` كعبارة وداع من البوت.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل في وحدة التحكم الرسالة `"Hello! I'm your coding fun fact guide!"`.

```js
assert.equal(getLogs()[0], "Hello! I'm your coding fun fact guide!")
```

يجب أن تعلن عن متغير `botName`. تحقق من أي أخطاء في التهجئة أو الحروف الكبيرة والصغيرة.

```js
assert.exists(botName);
```

يجب أن يكون المتغير `botName` سلسلة نصية.

```js
assert.isString(botName);
```

يجب أن تعلن عن متغير `botLocation`. تحقق من أي أخطاء في التهجئة أو الحروف الكبيرة والصغيرة.

```js
assert.exists(botLocation);
```

يجب أن يكون المتغير `botLocation` سلسلة نصية.

```js
assert.isString(botLocation);
```

يجب أن تعلن عن متغير `favoriteLanguage`. تحقق من أي أخطاء في التهجئة أو الحروف الكبيرة والصغيرة.

```js
assert.exists(favoriteLanguage);
```

يجب أن يكون المتغير `favoriteLanguage` سلسلة نصية.

```js
assert.isString(favoriteLanguage);
```

يجب أن تسجل في وحدة التحكم الرسالة `"My name is (botName) and I live on (botLocation)."` مع إضافة المتغيرات إلى السلسلة النصية.

```js
assert.equal(getLogs()[1], `My name is ${botName} and I live on ${botLocation}.`)
```

يجب أن تسجل في وحدة التحكم الرسالة `"My favorite programming language is (favoriteLanguage)."` مع إضافة المتغير إلى السلسلة النصية.

```js
assert.equal(getLogs()[2], `My favorite programming language is ${favoriteLanguage}.`)
```

يجب أن تستخدم `let` لإعلان متغير جديد باسم `codingFact`.

```js
const explorer = await __helpers.Explorer(code);
const { codingFact } = explorer.variables;
assert.isTrue(codingFact?.matches(`let codingFact = ${codingFact?.value?.toString()}`));
```

يجب أن تعطي المتغير `codingFact` قيمة تتضمن المتغير `favoriteLanguage`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.codingFact);
assert.include(String(codingFact), favoriteLanguage);
```

يجب أن تسجل المتغير `codingFact` في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.include(getLogs()[3], favoriteLanguage);
assert.isAtLeast(loggingCodingFacts.length, 1);
```

يجب أن تعيّن قيمة جديدة للمتغير `codingFact` تحتوي أيضًا على `favoriteLanguage`، وتسجلها في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.exists(loggingCodingFacts);
assert.isAtLeast(loggingCodingFacts.length, 2);
assert.include(getLogs()[4], favoriteLanguage);
assert.notEqual(getLogs()[4], getLogs()[3]);
```

يجب أن تعيّن قيمة للمتغير `codingFact` للمرة الثالثة تحتوي أيضًا على `favoriteLanguage`، وتسجلها في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingCodingFacts = codeWithoutComments.match(/console\.log\(\s*codingFact\s*\)/g)
assert.lengthOf(loggingCodingFacts, 3);
assert.notEqual(getLogs()[5], getLogs()[4]);
assert.isNotEmpty(codingFact);
```

يجب أن تسجل في وحدة التحكم الرسالة `"It was fun sharing these facts with you. Goodbye! - (botName) from (botLocation)."` مع إضافة قيم المتغيرات.

```js
assert.equal(getLogs()[6], `It was fun sharing these facts with you. Goodbye! - ${botName} from ${botLocation}.`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
console.log("Hello! I'm your coding fun fact guide!");

const botName = "JsBot";
const botLocation = "planet Eris";
const favoriteLanguage = "JavaScript";

console.log("My name is " + botName + " and I live on " + botLocation + ".");
console.log("My favorite programming language is " + favoriteLanguage + ".");

let codingFact = "Did you know that " + favoriteLanguage + " was created in just 10 days?";

console.log(codingFact);

codingFact = "Another fun fact: " + favoriteLanguage + "was originally called Mocha!";

console.log(codingFact);

codingFact = "Also, " + favoriteLanguage + " is the most popular programming language in the world.";

console.log(codingFact);

console.log("It was fun sharing these facts with you. Goodbye! - " + botName + " from " + botLocation + ".");

```
