---
id: 660ef55dd468079679ee0092
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Nun bist du bereit, deine nächste Variable zu erstellen. Entferne sowohl die `console.log` Anweisungen als auch die `character` Zuordnung.

Entferne auch deine `secondCharacter` Variable, aber verändere nicht die `character` Initialisierung.

# --hints--

Du solltest keine `console.log()`-Anweisungen in deinem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Du solltest den Wert der `character`-Variable neu zuweisen.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Deine `character` Variable sollte den Wert `"Hello"` haben.

```js
assert.equal(character, "Hello");
```

Du solltest keine `secondCharacter`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
