---
id: 660ef31a5be625914a0102cd
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Wenn eine Variable mit dem `let`-Schlüsselwort zugewiesen wird, kann man später den Wert <dfn>zuordnen</dfn> (oder den Wert ändern). In diesem Beispiel wird der Wert von `programmer` von `"Naomi"` zu `"CamperChan"` geändert.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Beachte: Bei der Neuzuweisung wird **NICHT** mehr das Schlüsselwort `let` verwendet.

Gib unterhalb deines `console.log` deiner `character` Variable den Wert `"World"`.

# --hints--

Du solltest `character` zweimal in deinem Code verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Du solltest den Zuweisungsoperator verwenden, um `character` neu zu benennen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Du solltest den String `"World"` in deinem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Deine `character` Variable sollte den Wert `"World"` haben.

```js
assert.equal(character, "World");
```

Deine neue Zuweisung sollte nicht `let` verwenden.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
