---
id: 660ef5105b8ba095307a0e50
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Gib jetzt deine `character`-Variable wieder in der Konsole aus. Du solltest in der Konsole den String `"Hello"` und dann den String `"World"` sehen.

# --hints--

Du solltest `console.log` ein zweites Mal verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Du solltest `character` ein zweites Mal in der Konsole loggen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


Dein neuer `console.log()` sollte nach deiner Neuzuordnung sein.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
