---
id: 660ef857f2806aa626d29d17
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Du kannst auch Subtraktion (`-`), Multiplikation (`*`) und Division (`/`) durchführen. Experimentier doch einfach mal etwas mit den Operatoren und Zahlen in deiner `console.log`. Wenn du bereit bist, zum nächsten Schritt fortzufahren, entferne des `console.log`.

# --hints--

Du solltest keinen `console.log()` in deinem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Deine `character` Variable sollte immer noch den Wert `"Hello"` haben.

```js
assert.equal(character, "Hello");
```

Deine `count` Variable sollte immer noch den Wert `8` haben.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
