---
id: 660f4f79e2a82a4e92290f44
title: Schritt 110
challengeType: 1
dashedName: step-110
---

# --description--

Hin und wieder möchtest du vorherigen Code, den du auskommentiert hast zurückbringen. Du kannst dies tun, indem du die `/*` und `*/` um diesen Code herum entfernst. Dies heißt <dfn>einkommentieren</dfn>.

Kommentieren deine erste `for`Schleife wieder ein. Lass den einzelnen Kommentar und die anderen beiden mehrzeiligen Kommentare auskommentiert.

# --hints--

Du solltest deine erste `for` Schleife einkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Du solltest deinen einzeiligen Kommentar nicht entfernen.

```js
assert.match(code, /\/\//);
```

Du solltest deine `while` Schleife nicht einkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Du solltest deine zweite `for` Schleife nicht einkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
