---
id: 660f4de78f775e480ba2e451
title: Schritt 107
challengeType: 1
dashedName: step-107
---

# --description--

Die `. unshift()` Methode eines Arrays erlaubt das Hinzufügen eines Wertes zum **Anfang** des Arrays anders als `. push()`, der den Wert am Ende des Arrays hinzufügt. `.unshift()` gibt die neue Länge des Arrays zurück, auf dem es aufgerufen wurde.

```js
const countDown = [2, 1, 0];
const newLength = countDown.unshift(3);
console.log(countDown); // [3, 2, 1, 0]
console.log(newLength); // 4
```

Verwende `const` um eine `unshifted` Variable zu deklarieren und weise ihr das Ergebnis des Aufrufs `.unshift()` deines `numbers` Arrays. Übergib `5` als Argument. Drucke dann deine`unshifted`-Variable.

# --hints--

Du solltest deine `unshifted`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+unshifted/);
```

Du solltest die `.unshift()`-Methode an deinem `numbers`array verwenden.

```js
assert.match(__helpers.removeJSComments(code), /numbers\.unshift\(/);
```

Du solltest `5` als Argument an deinen `unshift()` Aufruf übergeben.

```js
assert.deepEqual(numbers, [5,1,2,3]);
```

Du solltest das Ergebnis deines `.unshift()` Aufrufs deiner `unshifted` Variable zuweisen.

```js
assert.equal(unshifted, 4);
```

Du solltest deine `unshifted`-Variable loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*unshifted\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];

console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
