---
id: 660f3dd626be3a1ffe27e5d1
title: Schitt 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript unterstützt auch mehrzeilige Kommentare. Ein mehrzeiliger Kommentar beginnt mit `/*` und endet mit `*/`.

Anders als ein einzeiliger Kommentar wird ein mehrzeiliger Kommentar mehrere Zeilen zusammenfassen.

Verwende `/*` und `*/`, um deine aktuelle `for` Schleife, einschließlich des Körpers, in einen mehrzeiligen Kommentar zu verwandeln.

# --hints--

Du solltest einen mehrzeiligen Kommentar mit `/*` starten.

```js
assert.match(code, /\/\*/);
```

Du solltest einen mehrzeiligen Kommentar mit `*/` beenden.

```js
assert.match(code, /\*\//);
```

Deine gesamte `for` Schleife sollte auskommentiert werden.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
