---
id: 660f17d4e9f227d86e834abd
title: Schritt 33
challengeType: 1
dashedName: step-33
---

# --description--

Um eine Pyramide zu erzeugen, musst Du mehrere Zeilen erstellen. Wenn Du eine Aufgabe wiederholt ausführen mussst, bis eine Bedingung erfüllt ist, wirst du eine <dfn>Schleife</dfn> verwenden. Es gibt viele Möglichkeiten, eine Schleife zu schreiben.

Du beginnst mit einer gewöhnlichen `for` Schleife. `for` Schleifen verwenden folgende Syntax:

```js
for (iterator; condition; iteration) {
  logic;
}
```

In den kommenden Schritten wirst du jede einzelne Komponente einer Schleife im Detail erkunden. Baue vorerst eine `for` Schleife, die die Begriffe `"iterator"`, `"condition"`und `"Iteration"` für die drei Komponenten enthält. Behalte den <dfn>Körper</dfn> der Schleife, den Abschnitt innerhalb der geschweiften Klammern `{}`, leer.


# --hints--

Dein code sollte eine `for`-Schleife haben.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Die erste Komponente deiner `for` Schleife sollte der String `"iterator"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Die zweite Komponente deiner `for` Schleife sollte der String `"condition"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Die dritte Komponente deiner `for` Schleife sollte der String `"iteration"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Das Innere deiner `for`-Schleife sollte leer sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
