---
id: 660f061d259bbebc37461080
title: Schritt 20
challengeType: 1
dashedName: step-20
---

# --description--

Du kannst auf die Werte innerhalb eines Array zugreifen, indem du den <dfn>index</dfn> des Wertes verwendest. Ein Index ist eine Zahl, die die Position des Wertes im Array repräsentiert, beginnend bei `0` für den ersten Wert.

Du kannst auf den Wert zugreifen, indem du die <dfn>Klammernschreibweise</dfn>, wie `array[0]` verwendest.

Verwende `console.log`und die Klammernschreibweise, um den ersten Wert in deinem `rows`Array auszugeben.

# --hints--

Du solltest eine `console.log()`-Anweisungen in deinem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Du solltest auf deinen `rows`Array zugreifen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Du solltest die Klammer-Schreibweise für den `rows` Array anwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Du solltest die Klammerschreibweise verwenden, um auf das erste Element in deinem `rows`-Array zuzugreifen.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Du solltest das erste Element deines `rows`Arrays ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
