---
id: 660f4774e3e0df35a68bb5f2
title: Schritt 90
challengeType: 1
dashedName: step-90
---

# --description--

Um deine Pyramide wieder erzeugen zu lassen, verwende das Ergebnis des Aufrufs von `padRow` mit `done` und `count` als Argumente im `rows` Array, aus so wie du es in der ersten Schleife getan hast.

# --hints--

Deine Schleife soll die `.push()` Methode in deinen `rows` aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

Du solltest in deiner `.push()`-Methode deine `padRow()` Funktion aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

Du solltest `done` als erstes Argument in deinem `padRow`-Aufruf haben.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

Du solltest `count` als zweites Argument an deinen `padRow` Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
