---
id: 660f46460f9c36330ebc07d8
title: Schritt 88
challengeType: 1
dashedName: step-88
---

# --description--

Der Gleichstellungsoperator kann zu einem seltsamen Verhalten in JavaScript führen. Zum Beispiel ist `"0" == 0` wahr, obwohl es sich beim einen um eine Zeichenkette und beim anderen um eine Zahl handelt.

Der <dfn>strenge Gleichheit</dfn> Operator (strict equality) `===` wird verwendet, um zu überprüfen, ob zwei Werte gleich sind und den gleichen Typ haben. In der Regel ist dies der Gleichheits-Operator, den du immer verwenden solltest. Mit dem strengen Gleichheitsoperator wird `"0" === 0` falsch, weil sie zwar den gleichen Wert von Null haben, aber nicht den gleichen Typ.

Aktualisiere deine `done == count` Bedingung, um den strengen Gleichstellungs-Operator zu verwenden.

# --hints--

Deine `if` Bedingung sollte die strenge Gleichheit verwenden.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
