---
id: 660f07d231941bc11719f664
title: Schritt 21
challengeType: 1
dashedName: step-21
---

# --description--

Arrays sind speziell, da sie als <dfn>veränderbar</dfn> betrachtet werden. Das bedeutet, dass Du den Wert an einem Index direkt ändern kannst.

Zum Beispiel würde dieser Code die Zahl `25` dem zweiten Element im Array zuweisen:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Aktualisiere das **dritte** Element deines `rows` Array mit der Zahl `10`. Drucke dann den `rows` Array auf deine Konsole.

# --hints--

Du solltest die Klammer-Notation auf den `rows` Array erneut anwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Du solltest das dritte Element des `rows`Arrays aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Du solltest den Zuweisungsoperator auf das dritte Element des `rows` Arrays verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Du solltest den Wert `10` dem dritten Element deines `rows` Array zuweisen.

```js
assert.equal(rows[2], 10);
```

Du solltest eine zweite `console.log()`-Anweisungen in deinem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Du solltest deinen `rows`-Array loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
