---
id: 660f2eccfe3f820304af1b39
title: Schitt 64
challengeType: 1
dashedName: step-64
---

# --description--

Um zu wissen, wie eine Zeile formatiert werden soll, muss deine `PadRow` Funktion wissen, auf welcher Zeilennummer Du dich befindest und wie viele Datensätze insgesamt generiert werden.

Der beste Weg, dies zu tun, ist das Erstellen von Funktionsparametern. Gib deiner `PadRow` Funktion eine `rowNumber` und `rowCount` an. Mehrere Parameter werden durch Komma getrennt:

```js
function name(first, second) {

}
```

# --hints--

Deine `padRow`-Funktion sollte eine `rowNumber`-enthalten.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Du brauchst ein Komma nach deinem `rowNumber` Parameter.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Deine `padRow`-Funktion sollte ein `rowCount`-enthalten.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
