---
id: 660f4e74f7fd3f4a99ac2e50
title: Schritt 108
challengeType: 1
dashedName: step-108
---

# --description--

Arrays verfügen auch über eine `.shift()`-Methode. Dadurch wird das **first** Element des Arrays entfernt, im Gegensatz zu `.pop()`, das das letzte Element entfernt. Hier ist ein Beispiel für die `.shift()` Methode:

```js
const numbers = [1, 2, 3];
numbers.shift();
```

Der `numbers` Array wäre `[2, 3]`.

Direkt unter den `numbers` Array, deklariere eine `shifted` Variable und weise ihr das Ergebnis des Aufrufs `.shift()` im `numbers` Array zu. In der nächsten Zeile protokolliere die `Variable` auf der Konsole.

# --hints--

Du solltest `const` verwenden, um eine `shifted`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+shifted/);
```

Du solltest den `.shift()`-Aufruf an deinem `numbers`array verwenden.

```js
assert.deepEqual(numbers, [5, 2, 3]);
```

Du solltest das Ergebnis deines `.shift()` Aufrufs deiner `shifted` Variable zuweisen.

```js
assert.equal(shifted, 1);
```

Du solltest deine `shifted`-Variable loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*shifted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];


const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
