---
id: 660f033cf051ebb50ea3bf48
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

Beim Programmieren müssen Sie oft mit vielen Daten arbeiten. Es gibt viele Datenstrukturen, die dir dabei helfen können, Daten zu organisieren und zu verwalten. Eine der grundlegendsten Datenstrukturen ist ein <dfn>Array</dfn>.

Ein <dfn>Array</dfn> ist ein nicht-primitiver Datentyp, der eine Reihe von Werten enthalten kann. Nicht primitive Datentypen unterscheiden sich von primitiven Datentypen, da sie komplexere Daten speichern können. Primitive Datentypen wie strings und numbers können immer nur einen Wert gleichzeitig enthalten.

Arrays werden mit eckigen Klammern (`[]` gekennzeichnet. Hier ist ein Beispiel für eine Variable die einen leeren Array als Wert hat:

```js
let array = [];
```

Deklariere eine `rows`-Variable und weise ihr ein leeres Array zu.

# --hints--

Du solltest eine `rows`-Variable haben.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Du solltest `let` verwenden, um deine `rows`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Du solltest deiner einen leeren Array `rows` zuweisen.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
