---
id: 660f18f059fe0fda192ce394
title: Schritt 34
challengeType: 1
dashedName: step-34
---

# --description--

Deine Schleife benötigt jetzt einen richtigen Iterator. Der <dfn>Iterator</dfn> ist eine Variable, die du spezifisch in deiner `for` Schleife deklarieren kannst, um zu bestimmen, wie die Schleife iteriert oder durch deine Logik geht.

Es ist ein gängiges Vorgehen `i` als Zähler-Variable in einer Schleife zu verwenden. Eine `for` Schleife erlaubt es, dies in den Klammern `()` festzulegen. Hier ist beispielsweise eine `for` Schleife, die eine `index-` Variable festlegt und ihr den Wert `100` zuweist.

```js
for (let index = 100; "second"; "third") {

}
```

Ersetzeden String `"iterator"` durch eine `let` Anweisung für die Variable `i`. Weise ihm den Wert `0` zum Start zu. Dies gibt der Variablen `i` den Wert `0` beim **ersten Lauf** der Schleife.

# --hints--

Du solltest `let` verwenden, um eine `i`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Du sollltest deiner `i`-Variablen `0` zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Deine `for`-Schleife sollte ausgeführt werden, wenn `i` größer als `0` ist.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
