---
id: 660f45ccf4ca5c31f253005a
title: Schritt 86
challengeType: 1
dashedName: step-86
---

# --description--

Wenn du jetzt dein `continueLoop` auf true änderst, wird deine `while`Schleife unendlich laufen. Dies nennt man eine <dfn>unendliche Schleife</dfn>, und du solltest darauf achten, so etwas zu vermeiden. Eine Endlosschleife kann dein System zum Absturz bringen und einen vollständigen Neustart benötige, um dies zu Beenden.

Um dies zu vermeiden, beginne damit den Inkrement-Operator zu verwenden, um den Wert der `done` Variable innerhalb der Schleife zu erhöhen.

# --hints--

Deine`while`Schleife sollte die `done`Variable erhöhen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
