---
id: 660f4efcb8068e4cb470dca1
title: Schritt 109
challengeType: 1
dashedName: step-109
---

# --description--

Nun, da du diese Methoden ausprobiert hast, kannst du einen anderen umgekehrten Pyramidenansatz durchführen. Aber zuerst musst du deine Experimente bereinigen.

Entferne deinen `numbers` Array und die Methodenaufrufe und die Aufrufprotokolle.

# --hints--

Du solltest keinen `numbers`-Array haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

Du solltest keine `unshifted`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

Du solltest keine `shifted`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

Du solltest kein `console.log`-Statement mehr haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
