---
id: 660f0f980e98e8cf77f1ce31
title: Schritt 30
challengeType: 1
dashedName: step-30
---

# --description--

Ändere deine `rows` Deklaration so, dass es wieder ein leeres Array ist.

Änder deinen `'Hello'` String so, dass doppelte Anführungszeichen verwendet werden. Generell spielt es keine Rolle, welche der beiden du bevorzugst, aber es ist besser, wenn man diese im ganzen Projekt einheitlich verwendet.

# --hints--

Dein `rows`-Array sollte leer sein.

```js
assert.empty(rows);
```

Der `"Hello"` String sollte doppelte Anführungszeichen verwenden.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
