---
id: 660f3915b41a441537ec9f5e
title: Schritt 72
challengeType: 1
dashedName: step-72
---

# --description--

Da du `i` nur um `1` erhöhst, kannst du den <dfn>increment operator</dfn> `++` verwenden. Dieser Operator erhöht den Wert einer Variable um 1, durch Aktualisierung der Zuweisung für diese Variable. Beispielsweise würde `test`hier `8` ergeben:

```js
let test = 7;
test++;
```

Ersetze deine Additions-Zuweisung durch den Increment-Operator für deine Schleifendurchführung.

# --hints--

Deine `for` Schleife sollte die Additions-Zuweisung mit `i` nicht verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Deine `for`-Schleife sollte den increment-Operator mit `i` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
