---
id: 660f0da9bf1035c9097af20a
title: Schritt 28
challengeType: 1
dashedName: step-28
---

# --description--

Du solltest `"freeCodeCamp"` in der Konsole sehen. Das liegt daran, dass `.pop()` den Wert zurückgibt, der aus dem Array entfernt wurde - und du hast vorher `"freeCodeCamp"` an das Ende des Arrays hinzugefügt.

Aber was gibt `.push()` zurück? Weise deiner vorhandenen `rows.push()` einer neuen `pushed` Variablen zu und gib diese in der Konsole aus.

# --hints--

Du solltest eine `pushed`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Du solltest `let` verwenden, um deine `pushed`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Du solltest `rows.push("freeCodeCamp")` deiner `pushed` Variablen zuweisen.

```js
assert.equal(pushed, 4);
```

Du solltest deine `pushed`-Variable loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
