---
id: 660f47afe4c98536715d5fa4
title: Schritt 91
challengeType: 1
dashedName: step-91
---

# --description--

Der <dfn>strenge Ungleichheit</dfn> Operator (strict inequality) `! =` erlaubt es dir zu überprüfen, ob zwei Werte nicht gleich sind oder nicht den gleichen Typ haben. Die Syntax ist ähnlich wie die des Gleichheitsoperator: `value !== 4`.

Aktualisiere die`while` Schleifenbedingung, um zu überprüfen, ob `done` nicht mit `count` überein stimmt.

# --hints--

Deine `while` Schleife sollte prüfen, ob `done` und `count` ungleich sind.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
