---
id: 660f50a21fe7645252804f2b
title: Schritt 112
challengeType: 1
dashedName: step-112
---

# --description--

Was wäre, wenn du eine Möglichkeit hättest, zwischen einer invertierten Pyramide und einer Standard-Pyramide umzuschalten?

Beginne indem du eine `invertierte` Variable deklarierst und ihr den Wert `true` zuweist. Du änderst diese Variable nicht im Code, aber du musst `let` verwenden, damit unsere Tests sie später ändern können.

# --hints--

Du solltest eine `inverted`Variable mit `let` deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted/);
```

Du solltest`inverted` mit dem Wert `true` initialisieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+inverted\s*=\s*true;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
--fcc-editable-region--

--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  rows.unshift(padRow(i, count));
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
