---
id: 660f4377a359972c521d3f4b
title: Schritt 79
challengeType: 1
dashedName: step-79
---

# --description--

Du wirst die Zeichenkette in der Konsole ausgegeben sehen da `true` tatsächlich wahr ist.

Ändere die Bedingung deines `if`Statements auf die boolsche `false`.

# --hints--

Deine `if`-Bedingung sollte `false` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if (true) {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
