---
id: 660f38c34a4de6141c0c369f
title: Schritt 71
challengeType: 1
dashedName: step-71
---

# --description--

Dein Pyramidengenerator funktioniert nun wie erwartet. Aber dies ist eine hervorragende Gelegenheit, den Code, den du geschrieben haben, weiter zu erkunden.

Der Additions Operator ist nicht die einzige Möglichkeit, Werte zu einer Variable hinzuzufügen. Der <dfn>Additions-Zuweisungs</dfn> Operator kann als Kurzform verwendet werden, um "verwende den ursprünglichen Wert der Variable, füge diesen Wert hinzu und weis das Ergebnis der Variable wieder zu." auszudrücken Zum Beispiel würden diese beiden Anweisungen dasselbe Ergebnis hervorbringen:

```js
test = test + 1;
test += 1;
```

Update your iteration statement in the `for` loop to use addition assignment.

# --hints--

Deine `for` Schleife sollte nicht `i = i +1` verwenden;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Deine `for` Schleife sollte die Additions-Zuweisung mit `i` zugewiesen werden.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
