---
id: 660f46b9c417a8341729a3ab
title: Schritt 89
challengeType: 1
dashedName: step-89
---

# --description--

Wenn `done` den Wert von `count` erreicht hat, wollen wir, dass die Schleife nicht mehr ausgeführt wird.

Innerhalb deines `if` Bodys, ordne den boolschen Wert `false` deiner `continueLoop` Variable zu.

# --hints--

Dein `if`Body sollte deiner `continueLoop`Variable den Wert `false` zuordnen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{\s*continueLoop\s*=\s*false;?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done === count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
