---
id: 660f4b641290da41b2cf0dd9
title: Schritt 103
challengeType: 1
dashedName: step-103
---

# --description--

Ebenso wie bei der Addition gibt es verschiedene Operatoren für die Subtraktion. Der <dfn>Subtraktionszuweisung</dfn> Operator `-=` subtrahiert den angegebenen Wert vom aktuellen Variablenwert dann wird das Ergebnis wieder der Variable zugewiesen.

Replace your iteration statement with the correct statement using the subtraction assignment operator.

# --hints--

Deine `for` Schleife sollte nicht `i = i - 1` verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

Deine `for` Schleife sollte die Subtraktionsanweisung verwenden, um `i` um `1` zu reduzieren.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
