---
id: 660f505d02b2bd513a1c3468
title: Schritt 111
challengeType: 1
dashedName: step-111
---

# --description--

Deine Pyramide steht nicht mehr auf dem Kopf. Dies liegt daran, dass Du neue Zeilen zum **Ende** des Arrays hinzufügst.

Aktualisiere den Schleifenkörper, um neue Zeilen am Anfang des Arrays hinzuzufügen.

# --hints--

Du solltest die `unshift`-Methode der `rows` verwenden.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Du solltest einen `padRow()` Aufruf als Argument für deine `.unshift()` Methode übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Du solltest `i` als erstes Argument in deinem `padRow()`-Aufruf haben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Du solltest `count` als zweites Argument an deinen `padRow()` Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
