---
id: 660f4934fb48f63abd5ae371
title: Schritt 95
challengeType: 1
dashedName: step-95
---

# --description--

Die Verwendung von `done` um die Anzahl der erzeugten Zeilen zu verfolgen, funktioniert, aber man kann die Logik tatsächlich ein bisschen weiter aufräumen.

Arrays haben eine eigene `length` Eigenschaft, mit der Du  siehst, wie viele Werte oder <dfn>Elemente</dfn>im Array sind. Du würdest auf diese Eigenschaft mit Syntax wie `myArray.length` zugreifen.

Beachte, dass `rows.length` im `padRow` Aufruf einen Off-by-one-Fehler ergeben würde, da `done` *vor* dem Aufruf erhöht wird.

Aktualisiere deine Bedingung, um zu überprüfen, ob `rows.length` kleiner als `count` ist.

# --hints--

Deine `while` Schleife sollte überprüfen, ob `rows.length` kleiner als `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
