---
id: 660f1b6e60bd9edf902c81fd
title: Schritt 36
challengeType: 1
dashedName: step-36
---

# --description--

Deine<dfn>Iteration</dfn> wird deiner Schleife sagen, was sie nach jedem Lauf mit dem Zähler machen soll.

Wenn du einer Variable einen neuen wert zuweist, kannst du vor der Zuweisung zum vorherigen Wert referenzieren. Damit kannst du zum Beispiel drei zu einer bestehenden Zahl hinzufügen. Beispielsweise würde `bees = bees + 3;` den Wert von `bees` um drei erhöhen.

Verwende diese Syntax, um deinen `"Iteration"` String durch eine Neuzuweisung zu ersetzen, die `i` um eins erhöht.

# --hints--

Du solltest deine `i`Variablen um eins erhöhen.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Du sollltest `i+1`- deiner `i`Variablen zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Deine `for` Schleife sollte `i` bei jeder Iteration um `1` erhöhen.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
