---
id: 660f09a2694b59c3a10ee304
title: Schritt 23
challengeType: 1
dashedName: step-23
---

# --description--

Entferne vorerst dein erstes Konsolen-Log und deine `rows[rows.length - 1]` Zuweisung. Behalte die zweite `rows` Log-Anweisung für später.

# --hints--

Du solltest deine `console.log(rows[0])`-Anweisung entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Du solltest deine `-rows[rows.length - 1]` Zuweisung entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Du solltest deine `console.log(rows])`-Anweisung nicht entfernen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
