---
id: 660f1a00ac619ddc1e259a66
title: Schritt 35
challengeType: 1
dashedName: step-35
---

# --description--

Die <dfn>-Bedingung</dfn> einer `for` Schleife gibt an, wie oft sie iterieren soll. Wenn die `condition` falsch wird, wird die Schleife beendet.

In JavaScript kann ein boolescher Wert entweder `true` oder `false` sein. Dies sind keine Strings - du lernst später mehr über den Unterschied.

Im Moment verwenden wir den <dfn>less then</dfn> Operator (`<`). Dies erlaubt dir zu überprüfen, ob der Wert auf der linken Seite kleiner ist als der Wert auf der rechten Seite. Zum Beispiel würde `count < 3` als `true` gelten, wenn `count` `2`ist, und `false` falls `count` `4`4.

Tausche deinen`"condition"`String mit einer Bedingung aus, die kontrolliert, ob `i` kleiner als `count`ist.

# --hints--

Du solltest den less then Operator verwenden.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Du solltest den less than Operator verwenden, um zu kontrollieren, ob `i` weniger als `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Deine `for` Schleife sollte `i< count` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
