---
id: 660f207334fabaeac3269c38
title: Schritt 41
challengeType: 1
dashedName: step-41
---

# --description--

Um den `result` String zu manipulieren, wirst du eine andere Art von Schleife verwenden. Im Besonderen eine `for...of` Schleife, die jedem Element in einem zählbaren Objekt zählt und vorübergehend einer Variable zuweist.

Die Syntax für eine `for...of` Schleife sieht wie folgt aus:

```js
for (const value of iterable) {

}
```

Beachte, dass Du `const` verwenden kannst, weil die Variable nur für eine einzelne Iteration existiert, nicht während der gesamten Schleife.

Erstelle eine `for...of` Schleife, die durch den `rows` Array iteriert, wobei jeder Wert einer `row` Variable zugewiesen wird.

# --hints--

Du solltest ein anderes `for` Keyword verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Du solltest eine `row`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Du solltest `const` verwenden, um deine `row`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Deine `for...of` Schleife sollte deine `row` Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Deine `row` Variable sollte aus `rows` extrahiert werden mit dem `of` Schlüsselwort.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Deine `for...of`-Schleife sollte leer sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
