---
id: 660f359af3e32e0f1a6880b7
title: Schritt 68
challengeType: 1
dashedName: step-68
---

# --description--

Du solltest nun die gleiche Anzahl von Zeichen in deiner Konsole sehen. Deine `padRow` Funktion macht genau das Gleiche, was das vorherige Programm gemacht hat, aber jetzt ist es in einem eigenen, wiederverwendbaren Abschnitt.

Verwende den Additionsoperator, um ein einzelnes Leerzeichen `" "` jeweils an den Anfang und das Ende deines wiederholten `character`-Zeichens zu verketten.

Denke daran, dass du den `+`-Operator verwenden kannst, um solche Zeichen zu verketten:

```js
" " + "string"
```

# --hints--

Du solltest ein einzelnes Leerzeichen an den Anfang des zurückgegebenen Wertes ketten.

```js
assert.match(padRow(1, 1), /^\s/);
```

Du solltest ein einzelnes Leerzeichen an das Ende des zurückgegebenen Wertes ketten.

```js
assert.match(padRow(1, 1), /\s$/);
```

Deine `padRow()`-Funktion sollte die wiederholte `character`Reihe mit einem Leerzeichen vor und nach der Reihe ausgeben.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
