---
id: 660f39b444fd6f16d1e49c1f
title: Schritt 73
challengeType: 1
dashedName: step-73
---

# --description--

Anstatt `i + 1` zu deinem `padRow` übergeben zu müssen, könntest du deine Schleife mit `1` starten. Dies würde es dir ermöglichen, eine Schleife mit einem Index zu erstellen.

Aktualisiere den Zähler für den Start auf `1` statt `0`.

# --hints--

Deine `for` Schleife sollte `i` bei `1` starten lassen.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
