---
id: 660f229d2dbe09ef2954a4a1
title: Schritt 42
challengeType: 1
dashedName: step-42
---

# --description--

Erinnere dich daran, dass du in deiner ersten Schliefe den Additions Operator `+` verwendet hast, um den Wert von `i` um `1` zu erhöhen.

Du kannst ähnlich bei einem string Wert vorgehen, indem d einen neuen string an einen bestehenden hängt. Beispielsweise würde `hello = hello + " World";` den String `" World"` zu dem bestehenden String hinzufügen, der in der `hello` Variable gespeichert ist. Dies nennt man <dfn>Verkettung</dfn>.

Benutze in deiner `for...of`-Schleife den Additionsoperator, um den Wert von `row` an den Wert von `result` zu hängen.

# --hints--

Du solltest den Verkettungsoperator auf deine `result` Variable anwenden.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Du solltest `row` an dein `result` hängen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Du solltest die gesamte Verkettung deiner `result-` Variable wieder zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
