---
id: 660f53ad3d39175c5d4335ac
title: Schritt 117
challengeType: 1
dashedName: step-117
---

# --description--

Dein Pyramidengenerator befindet sich jetzt in einem fertigen Zustand mit mehr Funktionalitäten, als du ursprünglich geplant hast! Der nächste Schritt ist, den Code zu bereinigen.

Entferne alle Kommentare, sowohl einzeilige als auch mehrzeilige, aus deinem Code.

# --hints--

Du solltest in deinem Code keinen einzeiligen Kommentar haben.

```js
assert.notMatch(code, /\/\//);
```

Du solltest in deinem Code keinen mehrzeiligen Kommentar haben.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Du solltest keine Kommentare in deinem Code haben.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Du solltest Code entfernen, der durch mehrzeilige Kommentare kommentiert wurde.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
