---
id: 660f0ee51d7460ce88cd248d
title: Schritt 29
challengeType: 1
dashedName: step-29
---

# --description--

Hast du erwartet, `4` in der Konsole zu sehen? `.push()` gibt die neue Länge des Arrays zurück, nachdem der angegebene Wert hinzugefügt wurde.

Es ist wichtig, sich darüber im Klaren zu sein, welche Werte eine Methode zurückgibt. Nimm dir etwas Zeit, um mit `.push()` und `.pop()` zu experimentieren. Wenn du soweit bist, entferne alle deine `.push()` und `.pop()` Aufrufe und deine `console.log` Anweisungen.

# --hints--

Sie sollten keinen `.push()` Aufruf haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Sie sollten keinen `.pop()` Aufruf haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Du solltest keine log-Statements haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Du solltest keine `popped`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Du solltest keine `pushed`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
