---
id: 660f415b76859a2736771607
title: Schritt 78
challengeType: 1
dashedName: step-78
---

# --description--

Deine Pyramide ist wieder verschwunden. Das ist in Ordnung - das war zu erwarten.

Bevor du deine neue Schleife erstellst, musst du etwas über `if` Anweisungen lernen. Eine <dfn>`if` Anweisung</dfn> erlaubt es dir, einen Code-Block nur auszuführen, wenn eine Bedingung erfüllt ist. Sie verwenden folgende Syntax:

```js
if (condition) {
  logic
}
```

Erstelle ein `if`-Anweisung mit dem booleschen`true` als Bedingung. Im Körper schreibe den String `"Condition is true"`.

# --hints--

Du solltest eine `if`-Anweisung erstellen.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Deine `if`-Anweisung sollte `true` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Dein `if` Körper sollte `"Condition is true"` protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
