---
id: 660f44f10ea40f300b896a5e
title: Schritt 84
challengeType: 1
dashedName: step-84
---

# --description--

Jetzt, nachdem du `if...else if...else` Anweisungen geübt hast, kannst du diese aus deinem Code entfernen.

Sobald du dies abgeschlossen hast, verwende `let` um eine `continueLoop` Variable zu deklarieren und ihr den boolean-Wert `false` zuzuweisen. Nutze dann `let`, um eine `done`Variable zu deklarieren und ihr den Wert `0` zuzuweisen.

# --hints--

Du solltest keine `if`-Anweisung haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /if\s*\(/);
```

Du solltest keine `else if`-Anweisung haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Du solltest keine `else`-Anweisung haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /else\s*\{/);
```

Du solltest `let` verwenden, um eine `continueLoop`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+continueLoop/);
```

Deine `continueLoop` Variable sollte den Wert `falsch` haben.

```js
assert.isFalse(continueLoop);
```

Du solltest `let` verwenden, um eine `done`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+done/);
```

Deine `done` Variable sollte den Wert `0` haben.

```js
assert.strictEqual(done, 0);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
} else {
  console.log("This is the else block");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
