---
id: 660f0a55847d6cc485f29ba5
title: Step 26
challengeType: 1
dashedName: step-26
---

# --description--

Eine <dfn>-Methode</dfn> in JavaScript ist eine Funktion, die mit bestimmten Werten oder Objekten verbunden ist. Ein Beispiel, das dir bereits begegnet ist, ist die `.log()` Methode, die Teil des `console`-Objekts ist.

Arrays haben ihre eigenen Methoden. Die Erste, die du kennenlernst, ist die `.push()` Methode. Diese erlaubt dir, einen Wert zum Ende eines Arrays hinzuzufügen. Hier ist ein Beispiel, um die Zahl `12` zum Ende eines Arrays hinzuzufügen:

```js
array.push(12);
```

Verwende `.push()` um den String `"freeCodeCamp"` zum Ende deines `row` Arrays hinzuzufügen. Füge diesen Code vor deiner `console.log` hinzu, um die Änderungen sehen zu können, die du an deinem Array vorgenommen hast.

# --hints--

Du solltest `.push()` in deinem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Du solltest die `.push()`-Methode deines `rows` Arrays verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Du solltest den String `"freeCodeCamp"` an deine `.push()` Methode übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
