---
id: 660f3ce51f70571e1c5227c8
title: Schritt 76
challengeType: 1
dashedName: step-76
---

# --description--

Kommentare können hilfreich sein, um zu erklären, warum dein Code einen bestimmten Ansatz verfolgt, oder um To-do-Notizen für sich selbst zu hinterlegen.

In JavaScript können Sie `//` verwenden, um einen einzeiligen Kommentar in Ihrem Code zu hinterlassen.

Add a single-line comment above your `for` loop to remind yourself to change the code to a different kind of loop.

# --hints--

Du solltest einen einzeiligen Kommentar mit `//` beginnen.

```js
assert.match(code, /\/\//);
```

Der Kommentar sollte mindestens fünf Zeichen lang sein.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
