---
id: 660f374d532dc41189cc9cc2
title: Schritt 69
challengeType: 1
dashedName: step-69
---

# --description--

Jetzt ist es an der Zeit für ein bisschen Mathe. Stell dir eine dreireihige Pyramide vor. Wenn wir diese zentriert darstellen wollen, würde sie etwas so aussehen:

```md
··#··
·###·
#####
```

Leerzeichen wurden durch Interpunktionen oder Mittelpunkte ersetzt, um eine bessere Lesbarkeit herzustellen. Wenn du das Muster extrapolierst, kannst du erkennen, dass die Leerzeichen am Anfang und am Ende der Zeile einem Muster folgen.

Aktualisiere deine Leerzeichenkette, so dass sie `rowCount -rowNumber` mal wiederholt wird.

Öffne die Konsole, um das Ergebnis zu sehen.

# --hints--

Du solltest `.repeat()` bei deinen `" "`Strings aufrufen, um sie `rowCount - rowNumber` mal zu wiederholen.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
