---
id: 660f3ba3cceef11b6ba08b59
title: Schitt 75
challengeType: 1
dashedName: step-75
---

# --description--

Leider ist jetzt der Boden der Pyramide verschwunden. Dies liegt daran, dass Du einen weiteren <dfn>off-by-one Fehler</dfn> erstellt hast.

Deine ursprüngliche Schleife ging für `i` Werte von `0` bis `7`, denn `count` ist `8 8` und Ihre Bedingung erfordert `i` weniger als `count`. Deine Schleife läuft jetzt für `i` Werte von `1` bis `7`.

Deine Schleife muss aktualisiert werden, um zu laufen, wenn `i` `8`ist. Wenn du dir deine Logik ansiehst, heißt das, dass deine Schleife laufen soll, solange `i`<dfn>kleiner oder gleich</dfn>`count` ist. Du kannst dazu den kleiner oder gleich Operator`<=`nutzen.

Aktualisiere deine Schleifenbedingung, um sie auszuführen, während `i` kleiner oder gleich `count` ist.

# --hints--

Deine `für` Schleife sollte nicht überprüfen, ob `i` weniger als `count` ist.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Deine `for` Schleife sollte überprüfen, ob `i` kleiner oder gleich `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
