---
id: 660f540c2176ea5dec01306d
title: Schritt 118
challengeType: 1
dashedName: step-118
---

# --description--

Gut gemacht! Experimentiere jetzt mit verschiedenen Werten f체r deine `character`, `count`und `inverted` -Variablen.

Wenn du bereit bist f체r das n채chste Projekt, dann setze `character` auf `"!"`, `count` auf `10`und `inverted` auf `false` um fortzufahren.

Herzlichen Gl체ckwunsch zum Abschluss des ersten JavaScript-Projekts!

# --hints--

Du solltest `character` auf `"!"` setzen.

```js
assert.equal(character, "!");
```

Du solltest `count` auf `10` setzen.

```js
assert.equal(count, 10);
```

Du solltest `inverted` auf `false` setzen.

```js
assert.equal(inverted, false);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "#";
const count = 8;
const rows = [];
let inverted = true;
--fcc-editable-region--

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

# --solutions--

```js
const character = "!";
const count = 10;
const rows = [];
let inverted = false;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
