---
id: 660f0c34aad72dc712b97624
title: Schritt 27
challengeType: 1
dashedName: step-27
---

# --description--

Eine andere Methode, die f端r dieses Projekt wichtig ist, ist die `.pop()` Methode. Sie entfernt das letzte Element aus einem Array und gibt das Element <dfn>zur端ck</dfn> (eng: returns).

Wenn eine Methode einen Wert zur端ckgibt, kannst du dir es so vorstellen, dass sie den Wert zu dir zur端ckgibt, sodass du ihn in anderen Teilen deines Codes verwenden kannst.

Erstelle eine neue Variable namens `popped` und weise ihr das Ergebnis von `rows.pop()` zu. Gib dann `popped` in der Konsole aus (console.log).

# --hints--

Du solltest eine Variable mit dem Namen `popped` deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Du solltest `let` verwenden, um eine Variable namens `popped` zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Du solltest die `.pop()`-Methode aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Du solltest die `.pop()` Methode in deinem `rows` Array aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Du solltest deine `popped`-Variable ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
