---
id: 660f2a70ad6225fa503e71c3
title: Schritt 46
challengeType: 1
dashedName: step-46
---

# --description--

Du bist auf dem richtigen Weg! An diesem Punkt stößt du auf einen sogemannten <dfn>off-by-one-Fehler</dfn>, ein häufiges Problem in null-basierten Indizierungssprachen wie JavaScript.

Der erste Index deines `rows` Arrays ist `0`, das ist der Grund, warum Du deine `for` Schleife mit `i = 0` starten. Aber null mal etwas zu wiederholen ergibt kein Ergebnis.

Um dies zu beheben, füge `1`zum Wert von `i` in deinem `repeat()`Aufruf hinzu. Weise es nicht wieder `i` zu, wie Du es in deinen Schleifenbedingungen getan hast.

# --hints--

Du solltest `1` zu `i` in deiner `.repeat()` Methode hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
