---
id: 660f4a1472f8e63d76162ce5
title: Schritt 99
challengeType: 1
dashedName: step-99
---

# --description--

Was, wenn du deine Pyramide auf den Kopf stellst oder <dfn>invertierst</dfn>? Zeit, es auszuprobieren!

Beginne, in dem du eine `for`-Schleife erstellst. Lege einen `i`-Variable an und setze den Wert auf `count`. Verwenden dann denn Boolean: `false` für deine Schleifenbedingung und die Iterations-Anweisung.

# --hints--

Dein code sollte eine `for`-Schleife haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Deine `for`-Schleife sollte eine `i`-Variable initialisieren und den Wert auf `count` setzen.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count/);
```

Deine `for`-Schleife sollte `false` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false/);
```


Deine `for`-Schleife sollte `false` als Iterations-Anweisung haben.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s*i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
