---
id: 660f280dda5040f707c76b4a
title: Schritt 45
challengeType: 1
dashedName: step-45
---

# --description--

Du hast jetzt eine Reihe von `#` Zeichen, aber die Pyramidenform fehlt noch. Glücklicherweise enthält die Variable `i` die aktuelle Zeilennummer in deiner Schleife, so dass du sie für die Erstellung einer Pyramidenstruktur verwenden kannst.

Um dies zu erreichen, verwende die `.repeat()` Methode, die für Strings verfügbar ist. Diese Methode akzeptiert eine Zahl als Argument, die die Anzahl der Wiederholungen des Zielzeichens angibt. Beispielsweise um mit `.repeat()` den String `"Code zu erstellen! Code! Code!"`:

```js
const activity = "Code! ";
activity.repeat(3);
```

Verwende die `.repeat()`Methode auf dein `character` und gib ihr die Anzahl über `i` an.

# --hints--

Du solltest die `.repeat()`-Methode verwenden.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Du solltest die `repeat()`-Methode an deiner `character`Variable anwenden.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Du solltest `i` in deiner `.repeat()` Methode mitgeben.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Du solltest die `.repeat()`-Methode in der `.push()`-Methode anwenden

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
