---
id: 660f48e1d3682f39e81843c4
title: Schritt 94
challengeType: 1
dashedName: step-94
---

# --description--

Dein Pyramidengenerator funktioniert noch immer. Trotzdem ist es möglich, wieder mit einer Endlosschleife zu enden.

Weil Du nur kontrollierst, ib `done`ungleich `count` ist, läuft die Schleife unendlich, falls `done` **größer** als `count`wäre.

Aktualisiere deine Schleifenbedingung, um zu kontrollieren, ob ` done ` kleiner oder gleich `count` ist.

# --hints--

Deine `while` Schleife sollte überprüfen, ob `done` kleiner oder gleich `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*done\s*<=\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
