---
id: 660f4cffb1459d45e34902d1
title: Schritt 106
challengeType: 1
dashedName: step-106
---

# --description--

Du kannst die umgekehrte Pyramide tats채chlich bauen, ohne "r체ckw채rts" Schleifen zu m체ssen, wie du es getan hast.

Um dies zu tun, musst du einige neue Arraymethoden lernen. Beginne damit, durch Verwenden von `const` eine `numbers`-Variable zu deklarieren. Weise ihr einen Array mit den Elemente `1`, `2`und `3` zu. Logge dann das `numbers` Array.

# --hints--

Du solltest `const` verwenden, um eine `numbers`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Deine `numbers`-Variable sollte ein Array sein.

```js
assert.isArray(numbers);
```

Deine `numbers` Array sollten die Elemente `1`, `2`und `3` in dieser Reihenfolge haben.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Du solltest deinen `numbers`-Array loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
