---
id: 660f23b53db70af0f2620e78
title: Schritt 43
challengeType: 1
dashedName: step-43
---

# --description--

Jetzt erscheinen alle Zahlen in der gleichen Zeile. Das passt so nicht, um eine Pyramide zu erstellen.

Du wirst eine neue Zeile zu jeder Reihe hinzufügen müssen. Allerdings führt das Drücken der Rückgabetaste, um einen Zeilenumbruch zwischen Anführungszeichen in JavaScript einzufügen, zu einem Parsing-Fehler. Stattdessen musst du die spezielle <dfn>Escape-Sequenz</dfn> `\n`verwenden, die als neue Zeile interpretiert wird, wenn die Zeichenkette protokolliert wird. Hier ein Beispiel:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Verwende einen zweiten Additionsoperator, um eine neue Zeile nach `result` und `row` hinzuzufügen.

# --hints--

Du solltest die `\n` Escape-Sequenz verwenden. Denk daran, dass es ein String sein muss, so dass es in Anführungszeichen stehen muss.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Du solltest `row` an dein `result` hängen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Du solltest die `\n` Escape-Sequenz mit deiner `row` Variable verketten.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Du solltest die gesamte Verkettung deiner `result-` Variable zuweisen.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
