---
id: 660f447efc0e722f016c1be0
title: Schritt 81
challengeType: 1
dashedName: step-81
---

# --description--

Der Text ist wieder aufgetaucht! Das liegt daran, dass `"false"` eine Zeichenkette ist, die bei der Auswertung zu einem booleschen `true` wird. Dies bedeutet, dass `"false"` ein wahrer Wert ist.

Ein <dfn>wahrer Wert</dfn> ist ein Wert, der als wahr gilt, wenn er nach booleanischer Art ausgewertet wird. Die meisten Werte, die dir in JavaScript begegnen, ergeben wahr.

Ein <dfn>falscher Wert</dfn> ist das Gegenteil dazu - ein Wert, der als falsch gilt, wenn er nach booleanisch ausgewertet wird. JavaScript hat eine definierte Liste von Falschwerten. Darunter sind `false`, `0`, `""`, `null`, `undefined`und `NaN`.

Ändere deine `if` Bedingung in einen leeren String `""`, was ein falscher Wert ist.

# --hints--

Deine `if`-Anweisung sollte `""` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
