---
id: 660f2b6fd54ac1fc142804dd
title: Schritt 47
challengeType: 1
dashedName: step-47
---

# --description--

Die Logik für die Formatierung dieser Pyramide wird wahrscheinlich kompliziert werden, was bedeutet, dass es an der Zeit ist, diesen Code in eine Funktion zu überführen.

Eine <dfn>-Funktion</dfn> ist ein Codeblock, der in deinen Anwendung wiederverwendet werden kann. Funktionen werden mit der folgenden Syntax deklariert:

```js
function name(parameter) {

}
```

Das Schlüsselwort `-function` gibt JavaScript an, dass die `name` Variable eine Funktion ist. `parameter` ist eine Variable, die einen Wert repräsentiert, der bei Verwendung an die Funktion übergeben wird. Eine Funktion kann so viele oder so wenige <dfn>parameter</dfn> haben, wie du möchtest. Wie bei einer `for` Schleife ist der Abstand zwischen den geschweiften Klammern der <dfn>function body</dfn>.

Lege eine `padRow` Funktion fest. Erstelle noch keine Parametervariablen. Die Funktion sollte leer sein. Denke daran die Hökerschreibweise (camel case) bei der Benennung anzuwenden.

# --hints--

Du solltest das `function` Keyword verwenden.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Du solltest eine `padRow`-Funktion deklarieren.

```js
assert.isFunction(padRow);
```

Deine `padRow()`-Funktion sollte keinen Parameter enthalten.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Deine `padRow()`-Funktion sollte ein leeres Body-Element haben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
