---
id: 660f383d4c772c12ff59904b
title: Schritt 70
challengeType: 1
dashedName: step-70
---

# --description--

Du kannst vollständige Ausdrücke als Argument übergeben. Die Funktion erhält das Ergebnis der Auswertung dieses Ausdrucks. Zum Beispiel würden diese beiden Anweisungen dasselbe Ergebnis hervorbringen:

```js
test(2 * 3 + 1);
test(7);
```

Wirf nochmal einen Blick auf das Muster:

```md
··#··
·###·
#####
```

Aktualisiere den `character` Wert, um ihn `2 * rowNumber - 1` Mal wiederholen zu können.

Öffne nochmal die Konsole, um das aktualisierte Ergebnis zu sehen.

# --hints--

Deine Funktion sollte `2* rowNumber -1` an deinen `.repeat()` Aufruf geben.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
