---
id: 660f2fbd45b520046cac68e8
title: Schitt 65
challengeType: 1
dashedName: step-65
---

# --description--

Denk daran, dass Du in einem vorhergehenden Schritt etwas zu Rückgabewerten erfahren hast. Eine Funktion kann einen Wert für deine Anmwendung <dfn>zurückgeben</dfn>.

In einer Funktion wird das `return` Keyword verwendet, um einen Rückgabewert zu spezifizieren. Beispielsweise würde diese Funktion den Wert zurückgeben, der dem ersten Parameter gegeben wurde:

```js
function name(parameter) {
  return parameter;
}
```

Benutze das Keyword `return`, um den Wert der `character` Variablen zurückzugeben, in `rowNumber` Wiederholungen.

# --hints--

Du solltest die `.repeat()`-Methode verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Du solltest die `repeat()`-Methode an deiner `character`Variable anwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Du solltest `owNumber` an deinen `.repeat()` Aufruf geben.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Du solltest das `return`-Schlüsselwort verwenden.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Du solltest das Ergebnis deines Aufrufs `.repeat()` zurückgeben.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
