---
id: 660f34e99571070d56d2f231
title: Schritt 67
challengeType: 1
dashedName: step-67
---

# --description--

Deine `padRow` Funktion hat zwei Parameter, die du festgelegt hast. Bei Aufruf einer Funktion werden diesen Parametern Werte zur Verfügung gestellt.

Die Werte, die du für einen Funktionsaufruf angibst, werden als <dfn>Argumente</dfn>bezeichnet, und du <dfn>übergibst</dfn> die Argumente an einen Funktionsaufruf. Hier ist ein Funktionsaufruf mit `"Hello"` als Argument:

```js
test("Hello");
```

Übergebe `i + 1` und `count` als Argumente an deinen `padRow` Aufruf. Wie Parameter werden auch Argumente durch ein Komma getrennt.

# --hints--

Du solltest `i+1 ` an `padRow()` übergeben.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Du brauchst ein Komma nach deinem `i+1` Parameter.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Du solltest `count` als zweites Argument übergeben.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
