---
id: 6610c77d50636722e5b6be17
title: Schritt 60
challengeType: 1
dashedName: step-60
---

# --description--

Unterhalb des ` return ` Statements, protokolliere den String `"This works!"` in der Konsole.

Nachdem du das gemacht hast, wirst du erkennen, dass der String `"That works!` nicht in der Konsole angezeigt wird und die `console.log("This works!")` Zeile grau hinterlegt ist.

Kopiere Sie das Konsolenprotokoll und füge es über der `return` Anweisung ein. Jetzt sollte der String `"This works!"` in der Konsole erscheinen.

Wichtig zu wissen über das`return`Keyword ist, dass es nicht nur einen Ausgabewert definiert, der durch dieFunktion ausgegeben werden soll, sondern es auch das Ausführen deines Codes innerhalb einer Funktion oder eines Block-Statements stoppt. Dies bedeutet, dass kein Code nach einem `return` Statement ausgeführt wird.

# --hints--

Deine `padRow` Funktion sollte die `test-` Variable zurückgeben.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Dein erstes`console.log` sollte nach deinem `return`Keyword kommen.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Dein zweites`console.log` sollte vor deinem `return`Keyword kommen.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";

  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
