---
id: 6610c538372aa61cc0f5b122
title: Schritt 57
challengeType: 1
dashedName: step-57
---

# --description--

Variablen in JavaScript sind in einem bestimmten <dfn>scope</dfn> verfügbar. Mit anderen Worten, wo eine Variable deklariert wird, bestimmt, wo sie in deinem Code verwendet werden kann.

Der erste Scope ist der globale Scope. Variablen, die außerhalb eines "Blocks" deklariert werden, wie eine Funktion oder `for` Schleife sind im <dfn>globalen Bereich</dfn>. Deine `-character`, `count`und `rows` Variablen sind alle im globalen Bereich.

Wenn eine Variable im globalen Anwendungsbereich liegt, kann eine Funktion per Definition darauf zugreifen. Hier ist ein Beispiel für eine Funktion mit einer globalen `-title` Variable:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Dieses Beispiel gibt `"Professor Naomi"` zurück. Aktualisiere deine `padRow` Funktion, um den Wert der Verkettung der `character` -Variable an den Anfang des `name` Parameters zu verketten.

# --hints--

Deine Funktion `padRow` sollte `character` an den Anfang von `name` hängen.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Deine `padRow` Funktion sollte das Ergebnis von `character + name` zurückgeben.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + "\n" + row;
}

console.log(result);
```
