---
id: 6610c1d97b1671140f95cfbb
title: Schritt 50
challengeType: 1
dashedName: step-50
---

# --description--

Füge nun eine Log-Anweisung hinzu, um den Wert deiner `call` Variable auszugeben.

# --hints--

Du solltest einen `console.log`-Aufruf hinzufügen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2)
```

Du solltest deine `call`-Variable ausgeben. Vergiss nicht das Semikolon.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*call\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
