---
id: 6610b741b54b90f0c0fb3d58
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Wenn Variablennamen aus mehr als einem Wort bestehen, gibt es spezifische Namenskonventionen für die Groß- und Kleinschreibung der Wörter. In JavaScript ist die Nutzungsvorgabe <dfn>camel case</dfn>.

Camel case bedeutet, dass das erste Word komplett klein geschrieben wird, aber die folgenden Worte jeweils mit großem Anfangsbuchstaben. Hier sind einige Beispiele der camel case Schreibweise:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Verwenden Sie camel case, um eine neue `secondCharacter` Variable zu deklarieren.

# --hints--

Du solltest eine `secondCharacter`-Variable haben.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Du solltest `let` verwenden, um eine Variable mit dem Namen `secondCharacter` zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Du solltest deiner `secondCharacter`-Variablen keinen Wert zuweisen. Vergiss nicht das Semikolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
