---
id: 6610c21b3ef82015573ffbbe
title: Schritt 51
challengeType: 1
dashedName: step-51
---

# --description--

Deine `call`-Variable hat einen `undefined` (undefinierten) Wert, obwohl du ihn definiert hast! Das liegt daran, dass deine `padRow`Funktion momentan keinen Wert liefert. Standardmäßig geben Funktionen `undefined` als Wert zurück.

Um etwas anderes zurückzugeben, musst du das `return` Schlüsselwort verwenden. Hier ist ein Beispiel für eine Funktion, die den String `"Functions are cool!"` zurückgibt:

```js
function demo() {
  return "Functions are cool!";
}
```

Verwende das Schlüsselwort `return` um die Funktion den String `"Hello!"` zurückgeben zu lassen.

# --hints--

Deine `padRow` Funktion sollte das `return` Schlüsselwort verwenden.

```js
assert.match(padRow.toString(), /return/);
```

Deine `padRow` Funktion sollte den String `"Hello!"-` zurückgeben.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
