---
id: 6610c8cfe4cf4d278e35c156
title: Schritt 63
challengeType: 1
dashedName: step-63
---

# --description--

Wie erwartet, gibt deine Funktion nun erneut `undefined` zurück. Deine `call`Variable wird nicht mehr benötigt, also entferne die `call`Deklaration und den `console.log`für die `call` Variable.

# --hints--

Du solltest keine `call` Deklaration haben.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Du solltest deine `call`-Variable nicht ausgeben.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
