---
id: 6610c87eac0f0b256d7b037e
title: Schitt 62
challengeType: 1
dashedName: step-62
---

# --description--

Da deine Funktion den Parameter nicht mehr verwendet hat, hat das Ändern des Arguments die Funktion nicht beeinflusst.

Entferne die `test` Deklaration und `return` Anweisung aus deiner `padRow` Funktion, damit die Funktion wieder leer ist.

# --hints--

Deine `padRow`-Funktion sollte keine `test`-Variable enthalten.

```js
assert.notMatch(padRow.toString(), /test/);
```

Deine `padRow` Funktion sollte keinen Wert zurückgeben.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
