---
id: 6610c16c4fa0df12c0e30675
title: Schritt 49
challengeType: 1
dashedName: step-49
---

# --description--

Du rufst deine `padRow` Funktion auf, machst aber nichts mit diesem Funktionsaufruf. Alle Funktionen in JavaScript geben einen Wert <dfn>zurück</dfn>, was bedeutet, dass sie das durch den Aufruf definierte Ergebnis liefern, damit du es an anderer Stelle verwenden kannst.

Um das Ergebnis des Aufrufs deiner `padRow` Funktion zu sehen, deklariere eine `call`-Variable und weise deinem bestehenden `padRow`Aufruf diese Variable zu.

# --hints--

You should declare a `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

You should use `const` to declare your `call` variable.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

You should assign `call` the result of your `padRow` call.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
