---
id: 664599653fcd6e97104f9261
title: Schritt 83
challengeType: 1
dashedName: step-83
---

# --description--

Manchmal möchte man einen anderen Code ausführen, wenn alle `if...else if`Bedingungen `false` sind. Du kannst dies tun, indem du einen `else` Block hinzufügst.

Ein `else` Block wird nur dann geprüft, wenn die Bedingungen in `if` und `else if` Blöcke nicht erfüllt sind.

Hier wird der `else` Block zum `else if` Block hinzugefügt.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Füge einen `else` Block zu deinem `else if` Block hinzu. Innerhalb des `else` Blocks logge den String `"This is the else block"` auf der Konsole.

Um die Ergebnisse in der Konsole zu sehen, kann man die `<` in der `else if` Anweisung manuell zu `>` ändern. Dies wird die Bedingung auf `false` setzen und der `else` Block wird ausgeführt.

# --hints--

Du solltest einen `else`-Block hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Innerhalb des `else` Blocks logge den String `"This is the else block"` auf der Konsole.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

