---
id: 664ed97c55a99e5ffab759dc
title: Schritt 13
challengeType: 0
dashedName: step-13
---

# --description--

Bevor wir fortfahren, solltest Du Dir einen Moment Zeit nehmen, um die Konzepte nochmals durchzugehen, die Du gelernt hast.

Verwende das Schlüsselwort `let` um eine `profession` Variable und eine `age` Variable zu deklarieren. Initialisiere `profession` mit dem String `"teacher"`, aber initialisiere nicht `age` mit irgendeinem Wert.

Logge beide Variablen in der Konsole, um die Ergebnisse zu sehen.

# --hints--

Du solltest eine `profession`-Variable deklarieren.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Du solltest `let` verwenden, um deine `profession`-Variable zu deklarieren.

```js
assert.match(code, /let\s+profession/);
```

Sie sollten den String `"teacher"` der `profession` - Variable zuweisen.

```js
assert.strictEqual(profession, "teacher");
```

Du solltest eine `age`-Variable deklarieren.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Du solltest `let` verwenden, um eine `age`-Variable zu deklarieren.

```js
assert.match(code, /let\s+age/);
```

Du solltest deiner `age`-Variable keinen Wert zuweisen.

```js
assert.isUndefined(age);
```

Du solltest `age` auf der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Du solltest `profession` auf der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
