---
id: 664ed9344ba3495e1f67f8ea
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Nachdem Du nun die Deklaration und Initialisierung überprüft hast, entferne den Code, den Du zur Überprüfung geschrieben hast.

Entferne nicht deine `character` - Variable.

# --hints--

Du solltest keine `profession`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Du solltest keine `age`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Du solltest keine `console.log`-Anweisungen für `age` haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Du solltest keine `console.log`-Anweisungen für `profession` haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Du solltest immer noch deine `character`-Variable haben.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
