---
id: 665732e62eb8b30c249b5c2c
title: Schritt 25
challengeType: 1
dashedName: step-25
---

# --description--

Nun bist du bereit, zum nächsten Satz von Array-Lektionen überzugehen.

Entferne deinen ganzen Code aus dem vorherigen Schritt.

# --hints--

Du solltest alle Verweise auf deine `cities` Variable entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
