---
id: 68e5293bd00d2fe134f58983
title: Paso 11
challengeType: 20
dashedName: step-11
---

# --description--

Ahora agrega un segundo método abstracto llamado `apply_discount` a tu clase `DiscountStrategy`. Usa el decorador `@abstractmethod` y dale los parámetros `self` y `product`.

Agrega indicaciones de tipo: `product` debe ser del tipo `Product`, y el método debe devolver un `float`. Usa `pass` en el cuerpo del método.

# --hints--

Tu clase `DiscountStrategy` debe tener un método llamado `apply_discount`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Tu método `apply_discount` debe tener el decorador `@abstractmethod`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Tu método `apply_discount` debe tener los parámetros `self` y `product: Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Tu método `apply_discount` debe tener una pista de tipo de retorno de `-> float`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
