---
id: 68e52994bc7ea2e3dec530a1
title: Pasos 23
challengeType: 20
dashedName: step-23
---

# --description--

Ahora necesitas una forma de gestionar múltiples estrategias de descuento y encontrar el mejor precio.

Crea una clase llamada `DiscountEngine`. Esta clase gestionará todas las estrategias de descuento y calculará el mejor precio para un producto.

Dale un método `__init__` que tome `self` y `strategies` como parámetros, donde `strategies` tiene una anotación de tipo `list[DiscountStrategy]`. Usar `list[DiscountStrategy]` especifica que el parámetro es una lista que contiene objetos `DiscountStrategy`. El tipo de retorno del método debe ser `None`. Dentro del método `__init__`, guarda `strategies` como un atributo de instancia.

# --hints--

Debes crear una clase llamada `DiscountEngine`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Tu clase `DiscountEngine` debe tener un método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Tu método `__init__` debe tener los parámetros `self` y `strategies: list[DiscountStrategy]`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Tu método `__init__` debe tener un tipo de retorno `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

Debes asignar `strategies` a `self.strategies` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
