---
id: 68e52994bc7ea2e3dec5309d
title: Paso 20
challengeType: 20
dashedName: step-20
---

# --description--

Ahora vamos a crear una tercera estrategia de descuento para usuarios premium. Crea una clase llamada `PremiumUserDiscount` que herede de `DiscountStrategy`.

A diferencia de las estrategias anteriores, esta no necesita un método `__init__` porque no almacena ninguna configuración. Aplicará un descuento fijo del 20% a todos los usuarios premium.

Por ahora, solo crea la clase vacía que hereda de `DiscountStrategy` usando `pass`.

# --hints--

Debes crear una clase llamada `PremiumUserDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

Tu clase `PremiumUserDiscount` debe heredar de `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
