---
id: 68e524c94e2e4ab7b4bfc7e2
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, vas a construir un calculador de descuentos que puede aplicar diferentes estrategias de descuento a productos. El sistema determinará el mejor precio para un cliente basado en múltiples reglas de descuento.

Comienza definiendo una clase llamada `Product`. Dale a tu clase un método `__init__` con los siguientes parámetros: `self`, `name` y `price`. Dentro del método `__init__`, asigna los parámetros `name` y `price` a atributos de instancia con el mismo nombre.

# --hints--

Debes crear una clase llamada `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Tu clase `Product` debe tener un método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Tu método `__init__` debe tener los siguientes parámetros: `self`, `name` y `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Debes asignar `name` a `self.name` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Debes asignar `price` a `self.price` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
