---
id: 68e5293bd00d2fe134f58984
title: Paso 12
challengeType: 20
dashedName: step-12
---

# --description--

Ahora es momento de crear una implementación concreta de tu clase abstracta. Crea una clase llamada `PercentageDiscount` que herede de `DiscountStrategy`.

Dale un método `__init__` que tome `self` y `percent` como parámetros. Dale una pista de tipo `int` a `percent` y establece el tipo de retorno del método en `None`. Dentro del método `__init__`, guarda `percent` como un atributo de la instancia.

# --hints--

Debes crear una clase llamada `PercentageDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Tu clase `PercentageDiscount` debe heredar de `DiscountStrategy`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Tu clase `PercentageDiscount` debe tener un método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Tu método `__init__` debe tener los parámetros `self` y `percent: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Tu método `__init__` debe tener una indicación de tipo de retorno de `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Debes asignar `percent` a `self.percent` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
