---
id: 68e5293bd00d2fe134f5898a
title: Paso 19
challengeType: 20
dashedName: step-19
---

# --description--

Probemos tu nueva estrategia de descuento. Después del código existente, crea una instancia de `FixedAmountDiscount` con un monto de `5` y asígnala a una variable llamada `fixed_discount`. Luego aplica este descuento al producto e imprime el resultado.

# --hints--

Debes crear una variable llamada `fixed_discount` con una instancia de `FixedAmountDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("fixed_discount")
    assert isinstance(fixed_discount, FixedAmountDiscount)
  `)
})
```

Tu instancia `FixedAmountDiscount` debe inicializarse con la cantidad `5`.

```js
({
  test: () => runPython(`assert fixed_discount.amount == 5`)
})
```

Debes llamar a `fixed_discount.apply_discount(product)` e imprimir el resultado.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_call("print(fixed_discount.apply_discount(product))")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
--fcc-editable-region--

--fcc-editable-region--
```
