---
id: 68e529226b0defdff26c4a8b
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora comenzarás a construir la interfaz del sistema de descuentos. Python proporciona una manera de definir clases abstractas que sirven como planos para otras clases. Esto se hace usando la clase `ABC` (Abstract Base Class) del módulo `abc`.

En la parte superior de tu código, agrega la siguiente declaración de importación:

```py
from abc import ABC, abstractmethod
```

Esto importa la clase `ABC` y el decorador `abstractmethod`, que usarás para definir métodos que deben ser implementados por cualquier clase que herede de tu clase abstracta.

# --hints--

Debes tener la declaración de importación mencionada en la parte superior de tu código.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
