---
id: 68e5293bd00d2fe134f58985
title: Paso 13
challengeType: 20
dashedName: step-13
---

# --description--

Dado que `PercentageDiscount` hereda de una clase abstracta, debes implementar todos los métodos abstractos. Comienza implementando el método `is_applicable`. El método `is_applicable` debe recibir como parámetros `self`, `product: Product` y `user_tier: str`. Debe tener una indicación de tipo de retorno `bool`.

Este método debe determinar si un descuento porcentual es válido. Devuelve `True` si `self.percent` es menor o igual a `70`, de lo contrario devuelve `False`. Esto previene descuentos poco realistas como un 90% de descuento.

# --hints--

Tu clase `PercentageDiscount` debe tener un método `is_applicable`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Tu método `is_applicable` debe tener los parámetros `self`, `product: Product` y `user_tier: str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Tu método `is_applicable` debería devolver `True` cuando `self.percent` sea menor o igual a `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Tu método `is_applicable` debería devolver `False` cuando `self.percent` sea mayor que `70`.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Tu método `is_applicable` debe tener una pista de tipo de retorno de `bool`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
