---
id: 68e529226b0defdff26c4a8a
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

Ahora, crea una instancia de `Product` con el nombre `Wireless Mouse` y un precio de `50.0`. Asígnala a una variable llamada `product`. Luego, imprime `product` en la consola para ver cómo se ve.

# --hints--

Debes tener una variable llamada `product`.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Debes asignar una instancia de `Product` a la variable `product`.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

Tu instancia `Product` debe tener el nombre `Wireless Mouse` y el precio `50.0`.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Debes imprimir `product` en la consola.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Cuando imprimas un `Product` con nombre `'Wireless Mouse'` y precio `50.0`, debería mostrar `Wireless Mouse - $50.0`.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
