---
id: 68e5293bd00d2fe134f58987
title: Paso 16
challengeType: 20
dashedName: step-16
---

# --description--

Ahora crearás otra estrategia de descuento. Crea una clase llamada `FixedAmountDiscount` que herede de `DiscountStrategy`.

Dale un método `__init__` que tome `self` y `amount` como parámetros. Dale al parámetro `amount` una pista de tipo `int` y una pista de tipo de retorno `None` para el método. Dentro del método `__init__`, guarda `amount` como un atributo de instancia.

# --hints--

Debes crear una clase llamada `FixedAmountDiscount`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

Tu clase `FixedAmountDiscount` debería heredar de `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

Tu clase `FixedAmountDiscount` debe tener un método `__init__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

Tu método `__init__` debe tener los parámetros `self` y `amount: int`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

Tu método `__init__` debe tener un tipo de retorno `None`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

Debes asignar `amount` a `self.amount` dentro de tu método `__init__`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
