---
id: 68e5293bd00d2fe134f58981
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Crea una clase llamada `DiscountStrategy` que herede de `ABC`. Esta será una clase base abstracta que define la interfaz para todas las estrategias de descuento.

Por ahora, solo usa `pass` en el cuerpo de la clase para crear una clase vacía.

# --hints--

Debes crear una clase llamada `DiscountStrategy`.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Tu clase `DiscountStrategy` debería heredar de `ABC`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
