---
id: 68e529226b0defdff26c4a89
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

En este momento, si creas una instancia de `Product` y la imprimes, verás una representación predeterminada del objeto que no es muy informativa. Por ejemplo, podría verse algo así como `<__main__.Product object at 0x7f9b8c2d1d30>`.

Para convertir la representación predeterminada del objeto en una cadena legible, crea un método `__str__` en tu clase `Product` que devuelva una cadena con el formato `name - $price` (donde `name` y `price` deben ser reemplazados por los valores de los atributos correspondientes).

# --hints--

Tu clase `Product` debería tener un método `__str__`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Tu método `__str__` debería devolver una cadena en el formato `name - $price`.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
