---
id: 68e52994bc7ea2e3dec530a4
title: Paso 27
challengeType: 20
dashedName: step-27
---

# --description--

Ahora crea un ciclo `for` que itere sobre `self.strategies`. Para cada `strategy`, verifica si es aplicable usando el método `is_applicable`. Si lo es, calcula el precio con descuento usando `apply_discount` y guárdalo en una variable llamada `discounted`. Luego, agrega el precio con descuento a la lista `prices`.

# --hints--

Debes crear un ciclo `for` que itere sobre `self.strategies`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_for_iter().is_equivalent("self.strategies")
  `)
})
```

Dentro del ciclo, tu sentencia `if` debe ser `if strategy.is_applicable(product, user_tier):`.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("strategy.is_applicable(product, user_tier)")
  `)
})
```

Si una estrategia es aplicable, debes llamar a `apply_discount` y agregar el resultado a `prices`.

```js
({
  test: () => runPython(`
assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_body().is_ordered("discounted = strategy.apply_discount(product)","prices.append(discounted)")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod
from typing import List

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: List[DiscountStrategy]) -> None:
        self.strategies = strategies

--fcc-editable-region--
    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]
        

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
