---
id: 68e79d91daf8fb4808f93c47
title: Paso 15
challengeType: 20
dashedName: step-15
---

# --description--

Después de implementar el método `apply_discount`, prueba tu estrategia de descuento. Crea una instancia de `PercentageDiscount` con `10` por ciento y asígnala a una variable llamada `discount`. Luego aplica el descuento usando `discount.apply_discount(product)`, e imprime el resultado para ver el descuento en acción.

# --hints--

Debes crear una variable llamada `discount` con una instancia de `PercentageDiscount`.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Tu instancia `PercentageDiscount` debe ser inicializada con un `10` por ciento.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Debes llamar a `discount.apply_discount(product)` e imprimir el resultado.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
