---
id: 68e9207ddcba053c5cb24413
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

En Python, las pistas de tipo son señales opcionales que indican a otros desarrolladores cuál se espera que sea el tipo de dato de una variable o función. Aquí tienes un ejemplo que añade la pista de tipo `str` al parámetro `name`.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Para los parámetros del método `__init__`, actualiza el parámetro `name` para que tenga la pista de tipo `str` y el parámetro `price` para que tenga la pista de tipo `float`.

# --hints--

Debes agregar la pista de tipo `str` al parámetro `name` y la pista de tipo `float` al parámetro `price` en el método `__init__` de la clase `Product`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
