---
id: 68f119ad905b29ff59d858c8
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

Al igual que en el paso anterior, en tu método `__str__`, agrega una pista de tipo de retorno de `str` ya que este método devuelve una representación en cadena del objeto.

# --hints--

Tu método `__str__` debe tener una indicación de tipo de retorno de `str`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
