---
id: 6140c9d35015ae0ba0c250e8
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`div` 要素の中に `body` を追加し、`class` の `wheel` を設定してください。

新しく作成した `div` の中に、`span` に設定された `class` を持つ `line` 要素を6つ、`div` に設定された `class` を持つ `cabin` 要素を6つ追加してください。

# --hints--

`div` 要素の中に新しい `body` を作成してください。

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert(divs?.length === 1);
```

`div` 要素の中の `body` は `class` の `wheel` を持つべきです。

```js
assert(document.querySelector('body div')?.classList?.contains('wheel'));
```

`.wheel` 要素の中に6つの `span` 要素を含めてください。

```js
assert(document.querySelectorAll('.wheel span')?.length === 6);
```

`span` 要素の中の6つの `.wheel` 要素は `class` の `line` を持つべきです。

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert(spans?.every(span => span?.classList?.contains('line')));
assert(document.querySelectorAll('.line')?.length === 6);
```

`.wheel` 要素の中に6つの `div` 要素を含めてください。

```js
assert(document.querySelectorAll('.wheel div')?.length === 6);
```

`div` 要素の中の6つの `.wheel` 要素は `class` の `cabin` を持つべきです。

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert(divs?.every(div => div?.classList?.contains('cabin')));
assert(document.querySelectorAll('.cabin')?.length === 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
