---
id: 6140c7e645d8e905819f1dd4
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

標準のボイラープレートから始めましょう。`DOCTYPE` 宣言、言語が英語に設定された `html` 要素、`head` と `body` 要素を追加してください。

正しい `meta` のための `charset` 要素、`title` 要素、そして `link` ファイル用の `./styles.css` 要素を追加してください。

`title` を `Ferris Wheel` に設定してください。

# --hints--

コードには `DOCTYPE` 宣言が含まれている必要があります。

```js
assert.match(code, /<!DOCTYPE/gi);
```

`DOCTYPE` 宣言の後にスペースを入れてください。

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

ドキュメントタイプは `html` に設定してください。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

`DOCTYPE` 宣言はタイプの後に `>` で閉じてください。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

`html` 要素は `lang` 属性が `en` の開始タグを持つ必要があります。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

`html` 要素は終了タグを持つ必要があります。

```js
assert.match(code, /<\/html\s*>/);
```

`DOCTYPE` 宣言は HTML の最初にある必要があります。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

開始 `head` タグが必要です。

```js
assert.match(code, /<head\s*>/i);
```

終了 `head` タグが必要です。

```js
assert.match(code, /<\/head\s*>/i);
```

開始 `body` タグが必要です。

```js
assert.match(code, /<body\s*>/i);
```

終了 `body` タグが必要です。

```js
assert.match(code, /<\/body\s*>/i);
```

`body` 要素は `head` 要素の直後の最初の要素である必要があります。

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

`head` と `body` 要素は兄弟要素である必要があります。

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

`head` 要素は `html` 要素の中の最初の要素である必要があります。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

コードには `meta` 要素が必要です。

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

`meta` 要素は `charset` 属性を持ち、その値は `UTF-8` である必要があります。

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

`meta` 要素は `head` 要素の中にある必要があります。

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

コードには `title` 要素が必要です。

```js
const title = document.querySelector('title');
assert.exists(title);
```

プロジェクトのタイトルは `Ferris Wheel` である必要があります。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

タイトルの大文字・小文字やスペルは正確にしてください。

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

コードには `link` 要素が必要です。

```js
assert.match(code, /<link/);
```

`link` 要素は 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素は `rel` 属性を持ち、その値は `stylesheet` である必要があります。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素は `href` 属性を持ち、その値は `styles.css` である必要があります。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
