---
id: 587d7faf367417b2b2512be9
title: JavaScript XMLHttpRequest 메서드로 데이터 전송하기
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

이전 예제에서는 외부 자원에서 데이터를 받았습니다. AJAX 요청을 지원하고 URL을 알고 있다면 외부 자원에 데이터를 보낼 수도 있습니다.

JavaScript의 `XMLHttpRequest` 메서드는 서버에 데이터를 전송할 때도 사용됩니다. 예시는 다음과 같습니다:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

이 메서드들을 이전에 여러 번 보셨을 겁니다. 여기서 `open` 메서드는 요청을 외부 자원의 주어진 URL에 대한 `POST`로 초기화하며, 세 번째 매개변수로 `true`를 전달해 비동기 작업임을 나타냅니다.

`setRequestHeader` 메서드는 요청자와 요청에 관한 정보를 담은 HTTP 요청 헤더의 값을 설정합니다. 이 메서드는 `open` 메서드 다음에, 그리고 `send` 메서드 전에 호출해야 합니다. 두 매개변수는 헤더 이름과 그 헤더 본문에 설정할 값입니다.

다음으로, `onreadystatechange` 이벤트 리스너는 요청 상태 변화를 처리합니다. `readyState`가 `4`이면 작업이 완료된 것이고, `status`가 `201`이면 요청이 성공한 것입니다. 따라서 문서의 HTML을 업데이트할 수 있습니다.

마지막으로, `send` 메서드는 `body` 값을 담아 요청을 전송합니다. `body`는 `userName`와 `suffix` 키로 구성되어 있습니다.

# --instructions--

코드를 수정해 API 엔드포인트에 `POST` 요청을 하세요. 그런 다음 입력 필드에 이름을 입력하고 `Send Message`를 클릭하세요. AJAX 함수는 `Reply from Server will be here.`를 서버에서 받은 데이터로 교체해야 합니다. 응답은 이름 뒤에 ` loves cats` 텍스트가 붙도록 형식을 맞추세요.

# --hints--

코드는 새로운 `XMLHttpRequest`를 생성해야 합니다.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

코드는 `open` 메서드를 사용해 서버에 `POST` 요청을 초기화해야 합니다.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

코드는 `setRequestHeader` 메서드를 사용해야 합니다.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

코드는 함수로 설정된 `onreadystatechange` 이벤트 핸들러를 가져야 합니다.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

코드는 `message` 클래스를 가진 요소를 가져와 `textContent`를 `userName loves cats`로 변경해야 합니다.

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

코드는 `send` 메서드를 사용해야 합니다.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
