---
id: 587d7faf367417b2b2512be8
title: 사용자의 GPS 좌표를 찾기 위한 위치 정보 가져오기
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

사용자의 현재 위치에 접근할 수 있는 또 다른 멋진 기능이 있습니다. 모든 브라우저에는 이 정보를 제공하는 내장된 navigator가 있습니다.

navigator는 사용자의 현재 경도와 위도를 가져옵니다.

이 사이트가 현재 위치를 알 수 있도록 허용할지 차단할지 묻는 프롬프트가 표시됩니다. 코드가 올바르기만 하면 도전 과제는 어느 쪽이든 완료할 수 있습니다.

허용을 선택하면 출력 폰트에 위도와 경도가 표시되는 것을 볼 수 있습니다.

다음은 이를 수행하는 코드입니다:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

먼저 `navigator.geolocation` 객체가 존재하는지 확인합니다. 존재하면 해당 객체의 `getCurrentPosition` 메서드를 호출하여 사용자의 위치를 비동기적으로 요청합니다. 요청이 성공하면 메서드 내 콜백 함수가 실행됩니다. 이 함수는 점 표기법을 사용해 `position` 객체의 위도와 경도 값을 접근하고 HTML을 업데이트합니다.

# --instructions--

사용자의 현재 위치를 확인하고 HTML에 삽입하기 위해 `script` 태그 안에 예제 코드를 추가하세요.

# --hints--

사용자의 현재 위치에 접근하려면 `navigator.geolocation`를 사용해야 합니다.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

사용자의 위도 위치를 표시하려면 `position.coords.latitude`를 사용해야 합니다.

```js
assert.match(code, /position\.coords\.latitude/g);
```

사용자의 경도 위치를 표시하려면 `position.coords.longitude`를 사용해야 합니다.

```js
assert.match(code, /position\.coords\.longitude/g);
```

사용자의 위치는 `div` 요소 내에 `id="data"`를 사용해 표시해야 합니다.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
