---
id: 587d7fae367417b2b2512be4
title: API에서 JSON 데이터 접근하기
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

이전 도전 과제에서 freeCodeCamp Cat Photo API에서 JSON 데이터를 가져오는 방법을 배웠습니다.

이제 반환된 데이터를 자세히 살펴보면서 JSON 형식을 더 잘 이해해 보겠습니다. JavaScript에서 사용하는 표기법을 다시 기억해 보세요:

<blockquote>[ ] -> 대괄호는 배열을 나타냅니다.<br>{ } -> 중괄호는 객체를 나타냅니다.<br>" " -> 큰따옴표는 문자열을 나타냅니다. JSON에서는 키 이름에도 사용됩니다.</blockquote>

API가 반환하는 데이터 구조를 이해하는 것은 필요한 값을 어떻게 가져올지에 영향을 주기 때문에 중요합니다.

오른쪽에서 `Get Message` 버튼을 클릭하여 freeCodeCamp Cat Photo API JSON을 HTML에 불러오세요.

JSON 데이터에서 처음과 마지막 문자는 대괄호 `[ ]`입니다. 이는 반환된 데이터가 배열임을 의미합니다. JSON 데이터의 두 번째 문자는 중괄호 `{`로, 객체의 시작을 나타냅니다. 자세히 보면 세 개의 별도 객체가 있음을 알 수 있습니다. JSON 데이터는 세 개의 객체로 이루어진 배열이며, 각 객체는 고양이 사진에 대한 정보를 담고 있습니다.

앞서 객체는 쉼표로 구분된 "키-값 쌍"을 포함한다고 배웠습니다. Cat Photo 예제에서 첫 번째 객체는 `"id":0`이며, 여기서 `id`는 키이고 `0`는 해당 값입니다. 마찬가지로 `imageLink`, `altText`, `codeNames`에 대한 키도 있습니다. 각 고양이 사진 객체는 같은 키를 가지지만 값은 다릅니다.

첫 번째 객체에서 또 다른 흥미로운 "키-값 쌍"은 `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`입니다. 여기서 `codeNames`은 키이고 값은 세 개의 문자열로 이루어진 배열입니다. 배열을 값으로 가지는 키뿐 아니라 객체 배열도 있을 수 있습니다.

배열과 객체에서 데이터를 접근하는 방법을 기억하세요. 배열은 대괄호 표기법을 사용해 특정 인덱스의 항목에 접근합니다. 객체는 대괄호 표기법이나 점 표기법을 사용해 특정 속성의 값을 가져옵니다. 다음은 첫 번째 고양이 사진의 `altText` 속성을 출력하는 예시입니다 — 편집기에서 파싱된 JSON 데이터는 `json`라는 변수에 저장되어 있습니다:

```js
console.log(json[0].altText);
```

콘솔에는 문자열 `A white cat wearing a green helmet shaped melon on its head.`이 출력됩니다.

# --instructions--

`id`가 2인 고양이에 대해 `codeNames` 배열의 두 번째 값을 콘솔에 출력하세요. 객체(`json` 변수에 저장됨)에서 대괄호 표기법과 점 표기법을 사용해 값을 접근해야 합니다.

# --hints--

코드는 대괄호 표기법과 점 표기법을 사용해 올바른 코드 이름에 접근하고 `Loki`을 콘솔에 출력해야 합니다.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
