---
id: 587d7fae367417b2b2512be5
title: JSON 데이터를 HTML로 변환하기
challengeType: 6
forumTopicId: 16807
dashedName: convert-json-data-to-html
---

# --description--

이제 JSON API에서 데이터를 가져왔으니, HTML에 표시할 수 있습니다.

고양이 사진 객체가 배열에 저장되어 있으므로 `forEach` 메서드를 사용해 데이터를 반복할 수 있습니다. 각 항목에 도달할 때마다 HTML 요소를 수정할 수 있습니다.

먼저 `let html = "";`으로 html 변수를 선언하세요.

그다음 JSON을 반복하면서 키 이름을 `strong` 태그로 감싸고 그 뒤에 값을 추가하는 HTML을 변수에 더하세요. 반복이 끝나면 렌더링합니다.

다음은 이 작업을 수행하는 코드입니다:

```js
let html = '';
json.forEach(function (val) {
  const keys = Object.keys(val);
  html += "<div class = 'cat'>";
  keys.forEach(function (key) {
    html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
  });
  html += '</div><br>';
});
```

**참고:** 이 도전 과제에서는 페이지에 새로운 HTML 요소를 추가해야 하므로 `textContent`에 의존할 수 없습니다. 대신 사이트가 교차 사이트 스크립팅 공격에 취약해질 수 있는 `innerHTML`를 사용해야 합니다.

# --instructions--

JSON 데이터를 반복하여 표시할 HTML 요소를 생성하는 `forEach` 메서드를 추가하세요.

다음은 예시 JSON입니다:

```json
[
  {
    "id": 0,
    "imageLink": "https://cdn.freecodecamp.org/curriculum/legacy-json-apis-ajax/funny-cat.jpg",
    "altText": "A white cat wearing a green helmet shaped melon on its head. ",
    "codeNames": ["Juggernaut", "Mrs. Wallace", "Buttercup"]
  }
]
```

# --hints--

데이터를 `html` 변수에 저장해야 합니다

```js
assert.match(__helpers.removeWhiteSpace(code), /html(\+=|=html\+)/g);
```

API에서 가져온 JSON 데이터를 반복하려면 `forEach` 메서드를 사용하세요.

```js
assert.match(code, /json\.forEach/g);
```

키 이름을 `strong` 태그로 감싸야 합니다.

```js
assert.match(code, /<strong>.+<\/strong>/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line

        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line
        json.forEach(function (val) {
          var keys = Object.keys(val);
          html += "<div class = 'cat'>";
          keys.forEach(function (key) {
            html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
          });
          html += '</div><br>';
        });
        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
