---
id: 587d7fad367417b2b2512be1
title: JavaScript의 onclick 속성을 사용해 클릭 이벤트 처리하기
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

페이지가 모두 로드된 후에만 코드가 실행되길 원합니다. 이를 위해 문서에 `DOMContentLoaded`라는 JavaScript 이벤트를 연결할 수 있습니다. 다음은 이를 구현한 코드입니다:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

`DOMContentLoaded` 함수 내부에 이벤트 핸들러를 구현할 수 있습니다. 사용자가 `onclick` 요소를 클릭할 때 실행되는 `#getMessage` 이벤트 핸들러를 다음 코드를 추가해 구현할 수 있습니다:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

id가 `DOMContentLoaded`인 요소에 대해 `getMessage` 함수 내부에 클릭 이벤트 핸들러를 추가하세요.

# --hints--

id가 `document.getElementById`인 요소를 선택하려면 `getMessage` 메서드를 사용하세요.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

코드에 `onclick` 이벤트 핸들러를 추가하세요.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
