---
id: 587d7fad367417b2b2512be2
title: 클릭 이벤트로 텍스트 변경
challengeType: 6
forumTopicId: 301500
dashedName: change-text-with-click-events
---

# --description--

클릭 이벤트가 발생하면 JavaScript를 사용해 HTML 요소를 업데이트할 수 있습니다.

예를 들어, 사용자가 `Get Message` 버튼을 클릭하면 `message` 클래스를 가진 요소의 텍스트가 `Here is the message`로 변경됩니다.

이것은 클릭 이벤트 내에 다음 코드를 추가해서 작동합니다:

```js
document.getElementsByClassName('message')[0].textContent =
  'Here is the message';
```

# --instructions--

`onclick` 이벤트 핸들러 안에 코드를 추가해 `message` 요소 안의 텍스트를 `Here is the message`로 변경하세요.

# --hints--

코드는 `document.getElementsByClassName` 메서드를 사용해 `message` 클래스를 가진 요소를 선택하고, 그 `textContent`을 주어진 문자열로 설정해야 합니다.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.textContent\s*?=\s*?('|")Here is the message\2/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      document.getElementsByClassName('message')[0].textContent =
        'Here is the message';
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
