---
id: 587d7fae367417b2b2512be3
title: JavaScript XMLHttpRequest 메서드로 JSON 가져오기
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

외부 소스에서 데이터를 요청할 수도 있습니다. 이때 API가 중요한 역할을 합니다.

API, 즉 애플리케이션 프로그래밍 인터페이스는 컴퓨터들이 서로 소통할 때 사용하는 도구임을 기억하세요. AJAX라는 기술을 사용해 API에서 받은 데이터로 HTML을 업데이트하는 방법을 배우게 됩니다.

대부분의 웹 API는 JSON이라는 형식으로 데이터를 전송합니다. JSON은 JavaScript 객체 표기법을 의미합니다.

JSON 구문은 JavaScript 객체 리터럴 표기법과 매우 비슷하게 생겼습니다. JSON은 `{`와 `}` 사이에 객체 프로퍼티와 그 현재 값들을 포함합니다.

이 프로퍼티와 값들은 흔히 "키-값 쌍"이라고 부릅니다.

하지만 API가 전송하는 JSON은 `bytes`로 보내지고, 애플리케이션은 `string`로 받습니다. 이들은 JavaScript 객체로 변환할 수 있지만 기본적으로는 JavaScript 객체가 아닙니다. `JSON.parse` 메서드는 문자열을 파싱하여 그 문자열이 설명하는 JavaScript 객체를 만듭니다.

freeCodeCamp의 Cat Photo API에서 JSON을 요청할 수 있습니다. 클릭 이벤트에 넣을 수 있는 코드는 다음과 같습니다:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

각 부분이 하는 일을 검토해 보겠습니다. JavaScript `XMLHttpRequest` 객체는 데이터를 전송하는 데 사용되는 여러 프로퍼티와 메서드를 가지고 있습니다. 먼저 `XMLHttpRequest` 객체의 인스턴스를 생성하여 `req` 변수에 저장합니다. 다음으로 `open` 메서드가 요청을 초기화합니다 — 이 예제는 API에서 데이터를 요청하므로 `GET` 요청입니다. `open`의 두 번째 인수는 데이터를 요청하는 API의 URL입니다. 세 번째 인수는 불리언 값으로, `true`이면 비동기 요청이 됩니다. `send` 메서드는 요청을 전송합니다. 마지막으로 `onload` 이벤트 핸들러가 반환된 데이터를 파싱하고 `JSON.stringify` 메서드를 적용해 JavaScript 객체를 문자열로 변환합니다. 이 문자열이 메시지 텍스트로 삽입됩니다.

# --instructions--

코드를 업데이트하여 freeCodeCamp Cat Photo API에 `GET` 요청을 생성하고 전송하세요. 그런 다음 `Get Message` 버튼을 클릭하세요. AJAX 함수가 `The message will go here` 텍스트를 API에서 받은 원시 JSON 출력으로 바꿉니다.

# --hints--

코드는 새로운 `XMLHttpRequest`를 생성해야 합니다.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

코드는 `open` 메서드를 사용해 freeCodeCamp Cat Photo API에 `GET` 요청을 초기화해야 합니다.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

코드는 `send` 메서드를 사용해 요청을 전송해야 합니다.

```js
assert.match(code, /\.send\(\s*\)/g);
```

코드는 `onload` 이벤트 핸들러를 함수로 설정해야 합니다.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

코드는 `JSON.parse` 메서드를 사용해 `responseText`를 파싱해야 합니다.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

코드는 `message` 클래스를 가진 요소를 찾아 그 내부 HTML을 JSON 데이터 문자열로 바꿔야 합니다.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
