---
id: 5ccfad82bb2dc6c965a848e5
title: JavaScript fetch 메서드로 JSON 가져오기
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

외부 데이터를 요청하는 또 다른 방법은 JavaScript `fetch()` 메서드를 사용하는 것입니다. 이는 `XMLHttpRequest`와 동일하지만 구문이 더 이해하기 쉽다고 여겨집니다.

다음은 `/json/cats.json`에 GET 요청을 하는 코드입니다.

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

참고: `fetch()` 메서드는 기본 `GET` 메서드로 `HTTP`를 사용합니다. 즉, 기본 데이터 조회를 위해 명시적으로 지정할 필요가 없습니다.

이 코드의 각 부분을 살펴보세요.

첫 번째 줄은 요청을 만드는 부분입니다. 그래서 `fetch(URL)`는 지정된 URL에 `GET` 요청을 만듭니다. 이 메서드는 프로미스를 반환합니다.

프로미스가 반환된 후 요청이 성공하면 `then` 메서드가 실행되어 응답을 받아 JSON 형식으로 변환합니다.

`then` 메서드도 프로미스를 반환하며, 다음 `then` 메서드가 이를 처리합니다. 두 번째 `then`의 인수는 찾고자 하는 JSON 객체입니다!

이제 `document.getElementById()`를 사용해 데이터를 받을 요소를 선택합니다. 그런 다음 요청에서 반환된 JSON 객체로 만든 문자열을 삽입하여 요소의 HTML 코드를 수정합니다.

# --instructions--

코드를 업데이트하여 `GET` 대신 `fetch` 메서드를 사용해 freeCodeCamp Cat Photo API에 `XMLHttpRequest` 요청을 생성하고 보내세요.

# --hints--

가져온 데이터를 사용해 내부 HTML을 교체해야 합니다.

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

`GET`를 사용해 `fetch` 요청을 해야 합니다.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

응답을 JSON으로 변환하려면 `then`를 사용하세요.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

다른 `then`가 JSON으로 변환한 데이터를 처리하려면 `then`를 사용하세요.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

id가 `message`인 요소를 가져와 JSON 데이터 문자열로 내부 HTML을 변경하세요.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
