---
id: 69899123889dcf5b19be1847
title: تنفيذ خوارزمية البحث بالعمق أولاً
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

في هذا المختبر، ستنفذ خوارزمية اجتياز رسم بياني تُسمى <dfn>البحث بالعمق أولاً</dfn> (DFS).

بينما يستكشف <dfn>البحث بالاتساع أولاً</dfn> (BFS) الجيران مستوى بمستوى، يغوص <dfn>البحث بالعمق أولاً</dfn> في عمق مسار واحد من الحواف قبل التراجع.

عندما تصل الخوارزمية إلى نهاية مسار (عقدة لا تحتوي على جيران غير مُزارين)، تتراجع إلى أحدث عقدة لا تزال تحتوي على حواف غير مستكشفة وتواصل البحث من هناك.

لتنفيذ هذه الخوارزمية بشكل تكراري، ستحتاج إلى استخدام **مكدس**. في JavaScript، تُستخدم `Array` القياسية كمكدس بشكل مثالي باستخدام طريقتي `.push()` و `.pop()`، متبعة مبدأ **آخر داخل، أول خارج (LIFO)**. هذا يضمن أن الجار الذي أُضيف مؤخرًا إلى المكدس هو التالي الذي يُستكشف.

مخرج بسيط لهذه الخوارزمية هو قائمة بالعقد التي يمكن الوصول إليها من عقدة معينة. لذلك، ستحتاج أيضًا إلى تتبع العقد التي تزورها.

**الهدف**: حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `dfs`.
1. يجب أن تأخذ دالة `dfs` معلمتين:

- `graph`: مصفوفة مجاورة (مصفوفة من المصفوفات).
    - `root`: مسمى رقمي لعقدة (فهرس البداية) وهو القيمة الرقمية للعقدة بين `0` و `n - 1`، حيث `n` هو العدد الكلي للعقد في الرسم البياني.

1. يجب أن تُرجع الدالة مصفوفة تحتوي على جميع مسميات العقد التي يمكن الوصول إليها من العقدة البداية.

# --hints--

يجب أن يكون لديك دالة باسم `dfs` تأخذ معلمتين.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` يجب أن تُرجع قائمة تحتوي على `1`، `2`، `3`، و `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` يجب أن تُرجع قائمة تحتوي على `1`، `2`، `3`، و `0`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` يجب أن تُرجع `[3]`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` يجب أن تُرجع قائمة تحتوي على `3` و `2`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` يجب أن تُرجع قائمة تحتوي على `0` و `1`.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
