---
id: 5e44412c903586ffb414c94c
title: 创建算术格式化器项目
challengeType: 23
saveSubmissionToDB: true
forumTopicId: 462359
dashedName: build-an-arithmetic-formatter-project
---

# --description--

小学生经常把算术题垂直排列，这样更容易解决。 例如，“235 + 52”变成：

```py
  235
+  52
-----
```

完成`algorithy_arranger`函数，该函数接收作为算术问题的字符串列表，并返回垂直并排排列的问题。 该函数应该接受可选的第二个参数。 当第二个参数设置为 `True` 时，应显示答案。

## 代码示例

函数调用：

```py
arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])
```

输出:

```py
   32      3801      45      123
+ 698    -    2    + 43    +  49
-----    ------    ----    -----
```

函数调用：

```py
arithmetic_arranger(["32 + 8", "1 - 3801", "9999 + 9999", "523 - 49"], True)
```

输出:

```py
  32         1      9999      523
+  8    - 3801    + 9999    -  49
----    ------    ------    -----
  40     -3800     19998      474
```

## 规则

如果提供的问题格式正确，该函数将返回正确的转换，否则，它将**返回**一个**字符串**描述对用户有意义的错误。

- 将返回错误的情况：
  - 如果向函数提供的<strong x-id=“1”>问题太多</strong>。 限制是 **五个**，任何更多的都会返回：
`'错误：问题太多。'`
  - 函数将接受的适当运算符是**加法**和 **减法**. 乘法和除法将返回错误。 本要点中未提及的其他操作员不需要进行测试。 本要点中未提及的其他操作符不需要进行测试。
  - 每个数字（操作数）只能包含数字。 否则，函数将返回：
`'Error: Numbers must only contain digits.'</代码>
  - 每个操作数（即运算符两侧的数字）的宽度最多为四位。 否则，返回的错误字符串将是:
`'Error: Numbers cannot be more than four digits.'`
- 如果用户提供了正确格式的问题，返回的转换将遵循以下规则：
  - 运算符和两个操作数中最长的操作数之间应该有一个空格，运算符将与第二个操作数在同一行，两个操作名将按照提供的顺序排列（第一个操作数是最上面的，第二个是最下面的）。
  - 数字应该右对齐。
  - 每个问题之间应该有四个空格。
  - 每个问题底部应有破折号。破折号应沿每个问题的整个长度单独运行。（上面的示例显示了这应是什么样子。）

注意：使用F12打开浏览器控制台，查看更详细的测试输出。

# --hints--

`arithmetic_arranger(["3801 - 2", "123 + 49"])` 应该返回 `  3801      123\n-    2    +  49\n------    -----`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3801 - 2", "123 + 49"]), '  3801      123\\n-    2    +  49\\n------    -----')`);
  }
})
```

`arithmetic_arranger(["1 + 2", "1 - 9380"])` 应该返回 `  1         1\n+ 2    - 9380\n---    ------`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["1 + 2", "1 - 9380"]), '  1         1\\n+ 2    - 9380\\n---    ------')`);
  }
})
```

`arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"])` 应该返回 `    3      3801      45      123\n+ 855    -    2    + 43    +  49\n-----    ------    ----    -----`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "3801 - 2", "45 + 43", "123 + 49"]), '    3      3801      45      123\\n+ 855    -    2    + 43    +  49\\n-----    ------    ----    -----')`);
  }
})
```

`arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"])` 应该返回 `  11      3801      1      123         1\n+  4    - 2999    + 2    +  49    - 9380\n----    ------    ---    -----    ------`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["11 + 4", "3801 - 2999", "1 + 2", "123 + 49", "1 - 9380"]), '  11      3801      1      123         1\\n+  4    - 2999    + 2    +  49    - 9380\\n----    ------    ---    -----    ------')`);
  }
})
```

`算术数组（[“44+815”、“909-2”、“45+43”、“123+49”、“888+40”、“653+87”]）`应返回`'Error: Too many problems.'`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["44 + 815", "909 - 2", "45 + 43", "123 + 49", "888 + 40", "653 + 87"]), 'Error: Too many problems.')`);
  }
})
```

`算术数组（[“3/855”，“3801-2”，“45+43”，“123+49”]）`应返回`"Error: Operator must be '+' or '-'."`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 / 855", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Operator must be '+' or '-'.")`);
  }
})
```

`算术数组（[“24+85215”，“3801-2”，“45+43”，“123+49”]）`应返回`'Error: Numbers cannot be more than four digits.'`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["24 + 85215", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers cannot be more than four digits.")`);
  }
})
```

`算术数组（[“98+3g5”，“3801-2”，“45+43”，“123+49”]）`应返回`'Error: Numbers must only contain digits.'`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["98 + 3g5", "3801 - 2", "45 + 43", "123 + 49"]), "Error: Numbers must only contain digits.")`);
  }
})
```

`arithmetic_arranger(["3 + 855", "988 + 40"], True)` 应该返回 `    3      988\n+ 855    +  40\n-----    -----\n  858     1028`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["3 + 855", "988 + 40"], True), "    3      988\\n+ 855    +  40\\n-----    -----\\n  858     1028")`);
  }
})
```

`arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True)` 应该返回 `   32         1      45      123      988\n- 698    - 3801    + 43    +  49    +  40\n-----    ------    ----    -----    -----\n -666     -3800      88      172     1028`。

```js
({
  test: () => {
    runPython(`
from unittest import TestCase

TestCase().assertEqual(arithmetic_arranger(["32 - 698", "1 - 3801", "45 + 43", "123 + 49", "988 + 40"], True), "   32         1      45      123      988\\n- 698    - 3801    + 43    +  49    +  40\\n-----    ------    ----    -----    -----\\n -666     -3800      88      172     1028")`);
  }
})
```

# --seed--

## --seed-contents--

```py
def arithmetic_arranger(problems, show_answers=False):

    return problems

print(f'\n{arithmetic_arranger(["32 + 698", "3801 - 2", "45 + 43", "123 + 49"])}')
```

# --solutions--

```py
def arithmetic_arranger(problems, result=False):

	lin1 = ""
	lin2 = ""
	lin3 = ""
	lin4 = ""

	if len(problems) > 5:
		return 'Error: Too many problems.'
	for problem in problems:
		[num1, sym, num2] = problem.split()
		sign = ['+', '-']
		if sym not in sign:
			return ("Error: Operator must be '+' or '-'.")
		if len(num1) > 4 or len(num2) > 4:
			return ("Error: Numbers cannot be more than four digits.")
		if not num1.isnumeric() or not num2.isnumeric():
			return ("Error: Numbers must only contain digits.")

		lin1 += "  " + num1 + "    " if len(num1) >= len(
		    num2) else " " * (len(num2) + 2 - len(num1)) + num1 + "    "
		lin2 += sym + " " + num2 + "    " if len(num2) >= len(
		    num1) else sym + " " * (len(num1) - len(num2) + 1) + num2 + "    "
		nmax = (len(num1) + 2) if len(num1) >= len(num2) else (len(num2) + 2)
		lin3 += "-" * nmax + "    "
		op = int(num1) + int(num2) if sym == "+" else int(num1) - int(num2)
		lin4 += (" " * (nmax - len(str(op)))) + str(op) + "    "

	arranged_problems = lin1.rstrip() + "\n" + lin2.rstrip(
	) + "\n" + lin3.rstrip()
	if result:
		arranged_problems += "\n" + lin4.rstrip()

	return arranged_problems
```
