---
id: 6564c67db26c417561ab510d
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Як видно з вихідних даних, тепер ваш регулярний вираз співпадає з першим `l` у рядку.

У вашому шаблоні ви можете додати квантор після символу, щоб вказати, скільки разів цей символ має повторюватися. Наприклад, квантор `+` означає, що символ має повторюватися один або більше разів.

Додайте квантор `+` до вашого шаблону.

# --hints--

Ви повинні змінити вашу змінну `pattern` на `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
