---
id: 656469e62377ab34a5bcb8d1
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Оголосіть дві нові змінні з іменами `digits` та `symbols` і призначте їм відповідно `string.digits` та `string.punctuation`.

# --hints--

Ви повинні оголосити змінну з іменем `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Ви повинні призначити `string.digits` вашій змінній `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Ви повинні оголосити змінну з іменем `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Ви повинні призначити `string.punctuation` вашій змінній `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
