---
id: 656485a8a3496d4a36b1496a
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Ці три змінні становлять можливі символи для вибору під час генерації пароля.

Безпосередньо перед ними додайте коментар `Define the possible characters for the password`.

# --hints--

Ви повинні додати коментар безпосередньо над змінною `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
