---
id: 65646ab6bced52360f44670f
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Тепер оголосіть змінну з ім’ям `all_characters` і призначте їй результат конкатенації ваших існуючих змінних.

# --hints--

Ви повинні оголосити змінну з ім’ям `all_characters`.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Ви повинні конкатенувати `letters`, `digits` і `symbols` та призначити результат своїй змінній `all_characters`.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
