---
id: 6564d97a5ef99783216a6229
title: Крок 44
challengeType: 20
dashedName: step-44
---

# --description--

Python надає особливий тип рядка, який називається *raw* рядок. Raw рядки мають префікс `r`. Головна відмінність від звичайних рядків полягає в тому, як вони обробляють символ зворотної скісної риски: у raw рядках зворотні скісні риски сприймаються як буквальні символи, а не як символи екранування. При написанні регулярних виразів використання raw рядків — це хороша практика, оскільки вони зазвичай містять багато символів `\`.

Перетворіть свій рядок `pattern` на raw рядок, додавши префікс `r`.

# --hints--

Ви маєте змінити свою змінну `pattern` на `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
