---
id: 656474314bf37d3c83a85143
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Функція `choice()` з модуля `random` приймає послідовність і повертає випадковий елемент цієї послідовності.

Змініть свій виклик `print()`, щоб використати функцію `choice()` і передати `all_characters` як аргумент.

# --hints--

Ви повинні змінити свій існуючий виклик `print()`, щоб вивести `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
