---
id: 6564642ce403fe29cc23fee7
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Модуль Python — це файл, що містить код, призначений для виконання певних завдань. Стандартна бібліотека Python містить багато модулів, які ви можете імпортувати та використовувати у своєму коді. Ви можете зробити це, використовуючи інструкцію `import`, за якою йде назва модуля.

Почніть цей проєкт із імпорту модуля `string`.

# --hints--

Вам слід імпортувати модуль `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
