---
id: 6564cbb6311a40783b3f5de6
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

Класи символів також дозволяють вказати діапазон символів для співпадіння. Потрібно вказати початковий і кінцевий символи, розділені дефісом, `-`. Символи йдуть у порядку Unicode.

Змініть вашу змінну `pattern`, щоб вона відповідала будь-якій літері `t`, перед якою стоїть літера в нижньому регістрі у змінній `quote`. Для цього використайте діапазон символів від `a` до `z`.

# --hints--

Вам слід змінити вашу змінну `pattern`, щоб вона відповідала будь-якій літері `t` у `quote`, перед якою стоїть літера в нижньому регістрі, використовуючи клас `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
