---
id: 65649122c7f77f519aaf0975
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Нарешті оголосіть змінну `new_password` і призначте їй результат виклику `generate_password`. Передайте `8` як аргумент у ваш виклик `generate_password`.

# --hints--

Ви повинні викликати `generate_password`, передаючи `8` як аргумент.

```js
({ test: () => assert.match(code, /generate_password\s*\(\s*8\s*\)/) })
```

Ви повинні призначити `generate_password(8)` змінній `new_password`.

```js
({
  test: () => {
    assert(__userGlobals.has("new_password"));
    assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*\)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--

--fcc-editable-region--
```
