---
id: 65649a631440b757f9f7b428
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

Регулярний вираз, або regex, — це шаблон, який використовується для співпадіння певної комбінації символів у рядку. Це дійсна альтернатива `if`/`else` умовним інструкціям, коли потрібно знайти або співпасти з шаблонами всередині рядка для цілей валідації, заміни символів та інших.

Функція `compile()` з модуля `re` компілює рядок, переданий як аргумент, у об’єкт регулярного виразу, який можна використовувати іншими методами `re`.

Оголосіть нову змінну `pattern` і призначте цій змінній значення `re.compile('i')`.

# --hints--

Ви повинні оголосити змінну `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Ви повинні призначити `re.compile('i')` вашій змінній `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
