---
id: 6564683821b2ee3174e7250e
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Ви можете отримати доступ до утиліт, визначених всередині імпортованого модуля, через точкову нотацію. Точкова нотація працює шляхом додавання крапки, за якою йде назва утиліти, до імені модуля. Наприклад, ось як отримати доступ до константи `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

У цьому проєкті ви будете використовувати різні константи з модуля `string`.

Оголосіть нову змінну з назвою `letters` і призначте їй `string.ascii_letters`.

# --hints--

Ви повинні оголосити змінну з назвою `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Ви повинні призначити `string.ascii_letters` вашій змінній `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
