---
id: 65646ffeaed2d238c562a014
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Загальною конвенцією є розміщення інструкцій `import` на початку вашого коду. І додатково, у випадку кількох інструкцій `import`, сортуйте їх в алфавітному порядку для покращення читабельності.

На початку вашого коду імпортуйте модуль `random`.

# --hints--

Ви повинні імпортувати модуль `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
