---
id: 6564883669b5af4b69f794cc
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ваша змінна `all_characters` — це рядок, утворений усіма літерами нижнього та верхнього регістру, усіма 10 цифрами та кількома спеціальними символами.

Безпосередньо перед нею додайте коментар із текстом `Combine all characters`.

# --hints--

Ви повинні додати коментар безпосередньо над вашою змінною `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
