---
id: 656ef89b4b486b7a16077864
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Кортеж — це ще одна вбудована структура даних у Python. Кортежі дуже схожі на списки, але вони визначаються за допомогою круглих дужок `()`, замість квадратних дужок. Також кортежі є незмінними, на відміну від списків.

```py
my_tuple = ('larch', 1, True)
```

Ваш список `constraints` буде зберігати кортежі. Першим елементом кожного кортежу буде параметр обмеження.

Змініть присвоєння списку `constraints`, додавши кортеж до свого списку. Використайте `nums` як перший елемент і порожній рядок як другий елемент.

# --hints--

Ви повинні додати кортеж до списку `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Ви повинні додати `(nums, '')` до списку `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
