---
id: 656ef54070c72161e6feb90f
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Здається, все гаразд, але було б добре мати спосіб перевірити, чи згенерований пароль відповідає певним функціям. Наприклад, мінімальній кількості спеціальних символів, цифр або літер у верхньому/нижньому регістрі. Ви дуже скоро цим займетесь.

Поки що перетворіть на коментарі останні дві лінії свого коду.

# --hints--

Вам слід перетворити на коментарі останні дві лінії свого коду.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
