---
id: 656f0d6bc60b58b9777a36c9
title: Крок 34
challengeType: 20
dashedName: step-34
---

# --description--

Клас символів позначається квадратними дужками і відповідає одному символу серед тих, що вказані між дужками. Наприклад, `ma[dnt]` відповідає або `mad`, або `man`, або `mat`.

Змініть свій шаблон так, щоб він відповідав `w`, за яким слідує або `h`, або `a`.

# --hints--

Ваш шаблон має використовувати клас символів, щоб відповідати `w`, за яким слідує або `h`, або `a`.

```js
({ test: () => assert(runPython(`
    pattern == "w[ha]" or pattern == "w[ah]"    
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'l+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
