---
id: 6577319039f4f7de9251b822
title: Крок 61
challengeType: 20
dashedName: step-61
---

# --description--

Замість використання циклу та змінної-фішки, ви можете досягти того самого результату іншим підходом, який ви реалізуєте в наступних кількох кроках.

`all()` — це вбудована Python функція, яка повертає `True`, якщо всі елементи в заданому ітерованому об’єкті оцінюються як `True`. Інакше вона повертає `False`.

Замініть свій існуючий `for` цикл і дві `if` інструкції однією `if` інструкцією. Для умови `if` використайте виклик функції `all()` і передайте порожній список як аргумент виклику функції.

# --hints--

Ви повинні замінити свій існуючий `for` цикл і дві `if` інструкції однією `if` інструкцією.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

Ваша нова умова `if` має бути `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

У тілі вашої нової `break` має бути `if`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
