---
id: 5cd9a70215d3c4e65518328f
title: Costruisci un conto alla rovescia
challengeType: 26
dashedName: build-a-countdown
---

# --description--

In questo laboratorio, costruirai una `countdown` funzione che prende un numero di partenza e genera una lista per il conto alla rovescia.

Quando chiami `countdown(n)`, dovrebbe restituire un array che inizia da `n` e conta alla rovescia di `1` fino a raggiungere `1`.

Per esempio:

- `countdown(5)` dovrebbe restituire `[5, 4, 3, 2, 1]`
- `countdown(1)` dovrebbe restituire `[1]`

Se il numero di partenza è inferiore a `1`, non c'è nulla da contare alla rovescia, quindi restituisci un array vuoto:

- `countdown(0)` dovrebbe restituire `[]`
- `countdown(-3)` dovrebbe restituire `[]`

Per completare il laboratorio, devi costruire il risultato usando la ricorsione e non devi usare loop (`for`, `while`), `Array.from()` o metodi di iterazione degli array (`forEach`, `map`, `filter`, `reduce`). Ogni chiamata dovrebbe produrre il proprio array risultato (non usare variabili globali per archiviare il conto alla rovescia).

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare una funzione chiamata `countdown`.
2. La funzione `countdown` deve prendere un singolo argomento `n` (un numero).
3. Se `n` è inferiore a `1`, `countdown` deve restituire un array vuoto.
4. Altrimenti, `countdown` deve restituire un array contenente i numeri interi da `n` fino a `1` in ordine decrescente.
5. Devi usare la ricorsione (la funzione deve chiamare se stessa) e non devi usare loop di alcun tipo (`for`, `while`, `Array.from()` o metodi di ordine superiore come `forEach`, `map`, `filter`, `reduce`).
6. Ogni chiamata ricorsiva deve usare un valore più piccolo della chiamata precedente (per esempio, `n - 1`) così la funzione raggiunge il suo caso base.
7. La soluzione non deve usare variabili globali per archiviare o memorizzare nella cache l'array risultato.
8. Chiamare `countdown` più volte con input diversi deve sempre restituire il risultato corretto e indipendente.

# --hints--

Devi avere una funzione chiamata `countdown`.

```js
assert.isFunction(countdown);
```

`countdown(-1)` deve restituire un array vuoto.

```js
assert.isEmpty(countdown(-1));
```

`countdown(10)` deve restituire `[10, 9, 8, 7, 6, 5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(10), [10, 9, 8, 7, 6, 5, 4, 3, 2, 1]);
```

`countdown(5)` deve restituire `[5, 4, 3, 2, 1]`

```js
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

Il tuo codice non deve basarsi su alcun tipo di loop (`for`, `while`, `Array.from()` o funzioni di ordine superiore come `forEach`, `map`, `filter` e `reduce`).

```js
const forEachSpy = __helpers.spyOn(Array.prototype, 'forEach');
const mapSpy = __helpers.spyOn(Array.prototype, 'map');
const filterSpy = __helpers.spyOn(Array.prototype, 'filter');
const reduceSpy = __helpers.spyOn(Array.prototype, 'reduce');
const fromSpy = __helpers.spyOn(Array, 'from');
try {
  countdown(5);
  assert.isEmpty(forEachSpy.calls);
  assert.isEmpty(mapSpy.calls);
  assert.isEmpty(filterSpy.calls);
  assert.isEmpty(reduceSpy.calls);
  assert.isEmpty(fromSpy.calls);
} finally {
  forEachSpy.restore();
  mapSpy.restore();
  filterSpy.restore();
  reduceSpy.restore();
  fromSpy.restore();
}
assert.notMatch(__helpers.removeJSComments(code), /\b(for|while)\b/);
```

Devi usare la ricorsione per risolvere questo problema.

```js
const explorer = await __helpers.Explorer(code);
const functionBody = __helpers.removeWhiteSpace(__helpers.removeJSComments(explorer.allFunctions.countdown?.toString()));
assert.isAtLeast([...functionBody.matchAll(/countdown\=?(?:function)?[\(\=]\S*?[\)\=]/g)].length, 2);
```

Non devono essere usate variabili globali per memorizzare l'array.

```js
countdown(1)
assert.deepStrictEqual(countdown(5), [5, 4, 3, 2, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function countdown(n) {
  if (n < 1) {
    return [];
  }

  return [n, ...countdown(n - 1)];
}
```
