---
id: a3bfc1673c0526e06d3ac698
title: 构建一个素数和计算器
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

在本实验中，你将构建一个计算器，该计算器接受一个数字并 `return` 小于或 `equal` 于该数字的所有素数的和。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `sumPrimes` 函数，它接受一个数字作为参数。
1. `sumPrimes` 函数应返回小于或等于所提供数字的所有素数之和。
1. 如果输入数字小于 `2`，函数应返回 `0`。

# --hints--

你应该有一个 `sumPrimes` 函数。

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` 应该返回 `17`。

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` 应该返回 `10`。

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` 应该返回 `2`。

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` 应该返回 `0`。

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` 应该返回 `73156`。

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
