---
id: 684aaf9ec670c68d20efd0d0
title: Jenga msimamizi wa usanidi wa mtumizi
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

Katika maabara hii, utajenga Msimamizi wa Usanidi wa Mtumizi unaomruhusu mtumizi kusimamia mipangilio yao kama vile mandhari, lugha, na arifa. Utatekeleza vitendakazi vya kuongeza, kusahihisha, kufuta, na kuangalia mipangilio ya mtumizi.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kufafanua kitendakazi kinachoitwa `add_setting` chenye vigezo viwili vinavyo wakilisha kamusi ya mipangilio na tuple inayojumuisha jozi ya ufunguo-thamani

1. Kitendakazi cha `add_setting` kinapaswa:
    
   - Badilisha ufunguo na thamani kuwa herufi ndogo.
   - Ikiwa mpangilio wa ufunguo upo, rudisha `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - Ikiwa mpangilio wa ufunguo haupo, ongeza jozi ya ufunguo-thamani kwenye kamusi iliyotolewa ya mipangilio na rudisha `Setting '[key]' added with value '[value]' successfully!`.
   - Ujumbe unaorudishwa unapaswa kuwa na ufunguo na thamani kwa herufi ndogo.

1. Unapaswa kufafanua kitendakazi kinachoitwa `update_setting` chenye vigezo viwili vinavyo wakilisha kamusi ya mipangilio na tuple inayojumuisha jozi ya ufunguo-thamani.

1. Kitendakazi cha `update_setting` kinapaswa:
       
   - Badilisha ufunguo na thamani kuwa herufi ndogo.
   - Ikiwa mpangilio wa ufunguo upo, sasisha thamani yake katika kamusi iliyotolewa ya mipangilio na rudisha: `Setting '[key]' updated to '[value]' successfully!`
   - Ikiwa mpangilio wa ufunguo haupo, rudisha `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - Ujumbe unaorudishwa unapaswa kuwa na ufunguo na thamani kwa herufi ndogo.


1. Unapaswa kufafanua kitendakazi kinachoitwa `delete_setting` chenye vigezo viwili vinavyo wakilisha kamusi ya mipangilio na ufunguo.

1. Kitendakazi cha `delete_setting` kinapaswa:
    
   - Badilisha ufunguo uliotumwa kuwa herufi ndogo.
   - Ikiwa mpangilio wa ufunguo upo, ondoa jozi ya ufunguo-thamani kutoka kamusi iliyotolewa ya mipangilio na rudisha `Setting '[key]' deleted successfully!`
   - Ikiwa mpangilio wa ufunguo haupo, rudisha `Setting not found!`
   - Ujumbe unaorudishwa unapaswa kuwa na ufunguo kwa herufi ndogo.

1. Unapaswa kufafanua kitendakazi kinachoitwa `view_settings` chenye kigezo kimoja kinacho wakilisha kamusi ya mipangilio.

1. Kitendakazi cha `view_settings` kinapaswa:
   
   - Rudisha `No settings available.` ikiwa kamusi iliyotolewa ya mipangilio ni tupu.
   - Ikiwa kamusi ina mipangilio yoyote, rudisha mfuatano wa herufi unaoonyesha mipangilio. Mfuatano huo unapaswa kuanza na `Current User Settings:` ikifuatiwa na jozi za ufunguo-thamani, kila moja kwenye mstari mpya na ufunguo ukiwa na herufi kubwa mwanzoni. Kwa mfano, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` inapaswa kurudisha:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Kwa ajili ya kupima msimbo, unapaswa kuunda kamusi inayoitwa `test_settings` kuhifadhi baadhi ya mapendeleo ya usanidi wa mtumizi.

# --hints--

Unapaswa kuunda kamusi inayoitwa `test_settings` na kuongeza baadhi ya thamani ndani yake.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Unapaswa kufafanua kitendakazi kinachoitwa `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

Kitendakazi cha `add_setting` kinapaswa kuwa na vigezo viwili.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` inapaswa kubadilisha ufunguo kuwa herufi ndogo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` inapaswa kubadilisha thamani kuwa herufi ndogo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` inapaswa kurudisha ujumbe wa kosa `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` inapaswa kuongeza jozi mpya ya ufunguo-thamani na kurudisha ujumbe wa mafanikio `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` inapaswa kuongeza kwa usahihi jozi ya ufunguo-thamani iliyotolewa kwenye kamusi.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Unapaswa kufafanua kitendakazi kinachoitwa `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

Kitendakazi cha `update_setting` kinapaswa kuwa na vigezo viwili.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

Kitendakazi cha `update_setting` kinapaswa kubadilisha ufunguo kuwa herufi ndogo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

Kitendakazi cha `update_setting` kinapaswa kubadilisha thamani kuwa herufi ndogo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` inapaswa kusasisha ufunguo uliopo na kurudisha ujumbe wa mafanikio `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` inapaswa kurudisha ujumbe wa kosa `Setting 'volume' does not exist! Cannot update a non-existing setting.` wakati ufunguo haupo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` inapaswa kusasisha kwa usahihi jozi ya ufunguo-thamani iliyotolewa kwenye kamusi.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Unapaswa kufafanua kitendakazi kinachoitwa `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

Kitendakazi cha `delete_setting` kinapaswa kuwa na vigezo viwili.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` inapaswa kubadilisha ufunguo kuwa herufi ndogo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` inapaswa kuondoa ufunguo uliopo na kurudisha ujumbe wa mafanikio `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` inapaswa kurudisha ujumbe wa kosa `Setting not found!` wakati ufunguo haupo.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` inapaswa kuondoa kwa usahihi ufunguo uliotolewa kutoka kwenye kamusi.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Unapaswa kufafanua kitendakazi kinachoitwa `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

Kitendakazi cha `view_settings` kinapaswa kuwa na kigezo kimoja.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` inapaswa kurudisha ujumbe `No settings available.` ikiwa kamusi iliyotolewa ni tupu.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` inapaswa kurudisha mipangilio iliyopangwa kwa kamusi isiyo tupu.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` inapaswa kuweka herufi kubwa mwanzoni mwa kila jina la mpangilio.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` inapaswa kuonyesha matokeo sahihi na kumalizika na tabia ya kuingia mstari mpya.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
