---
id: 6a89c69cbe1bb69845992d8b
title: Створіть генератор діапазону чисел
challengeType: 27
dashedName: build-a-range-of-numbers-generator
---

# --description--

У цій лабораторній роботі ви створите функцію `range_of_numbers`, яка використовує рекурсію для генерації списку чисел у заданому діапазоні.

Наприклад:

- `range_of_numbers(3, 9)` має повернути `[3, 4, 5, 6, 7, 8, 9]`.
- `range_of_numbers(5, 5)` має повернути `[5]`.

Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні створити функцію з назвою `range_of_numbers` з двома параметрами `start_num` та `end_num`.
1. Функція має повертати список послідовних цілих чисел, який починається з `start_num` і закінчується `end_num`, включно з обома значеннями.
1. Значення `start_num` завжди буде меншим або рівним `end_num`.
1. Ваша функція має використовувати рекурсію, викликаючи сама себе.
1. Ваша функція не повинна використовувати цикли `for` або `while`, генератори списків або функцію `range()`.
1. Коли `start_num` дорівнює `end_num`, функція має повертати список, що містить `start_num`. Це базовий випадок.
1. Для рекурсивного випадку функція має викликати сама себе з аргументом, який наближає до базового випадку, а потім додавати поточне число до повернутого списку.
1. Функція не повинна використовувати глобальні змінні для збереження повернутого списку.

# --hints--

Ви повинні мати функцію з назвою `range_of_numbers` з параметрами `start_num` та `end_num`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert not _func.is_empty()
assert _func.has_args("start_num, end_num")
`) })
```

`range_of_numbers(1, 5)` має повертати `[1, 2, 3, 4, 5]`.

```js
({ test: () => runPython(`assert range_of_numbers(1, 5) == [1, 2, 3, 4, 5]`) })
```

`range_of_numbers(6, 9)` має повертати `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`assert range_of_numbers(6, 9) == [6, 7, 8, 9]`) })
```

`range_of_numbers(4, 4)` має повертати `[4]`.

```js
({ test: () => runPython(`assert range_of_numbers(4, 4) == [4]`) })
```

`range_of_numbers(10, 15)` має повертати `[10, 11, 12, 13, 14, 15]`.

```js
({ test: () => runPython(`assert range_of_numbers(10, 15) == [10, 11, 12, 13, 14, 15]`) })
```

Ваша функція `range_of_numbers` має викликати сама себе.

```js
({ test: () => runPython(`
assert _Node(_code).block_has_call("range_of_numbers", "range_of_numbers")
`) })
```

Ваша функція `range_of_numbers` не повинна використовувати цикли, генератори списків або функцію `range()`.

```js
({ test: () => runPython(`
_func = _Node(_code).find_function("range_of_numbers")
assert len(_func.find_for_loops()) == 0
assert len(_func.find_whiles()) == 0
assert len(_func.find_comps()) == 0
assert len(_func.find_calls("range")) == 0
`) })
```

Виклик `range_of_numbers(1, 3)` не повинен впливати на наступний виклик. Після цього виклику `range_of_numbers(6, 9)` має і надалі повертати `[6, 7, 8, 9]`.

```js
({ test: () => runPython(`
range_of_numbers(1, 3)
assert range_of_numbers(6, 9) == [6, 7, 8, 9]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def range_of_numbers(start_num, end_num):
    if start_num == end_num:
        return [start_num]
    numbers = range_of_numbers(start_num, end_num - 1)
    numbers.append(end_num)
    return numbers
```
