---
id: 68de7360e1832ea180577a8a
title: Побудувати калькулятор n-го числа Фібоначчі
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `fibonacci`.
2. Ви повинні визначити список з назвою `sequence` всередині функції `fibonacci`, і він має бути ініціалізований значеннями `[0, 1]`.
3. Функція `fibonacci` має приймати один параметр — додатне ціле число `n`.
4. Виклик `fibonacci(n)` має використовувати підхід динамічного програмування для обчислення та повернення `n`-го числа з послідовності Фібоначчі, де кожне число є сумою двох попередніх.
5. Кожне обчислене число на позиції `n` у послідовності Фібоначчі має зберігатися у списку `sequence` за індексом `n - 1`.

# --hints--

Ви повинні мати функцію з назвою `fibonacci`.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

Ваша функція `fibonacci` має приймати один параметр.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Ви повинні мати список з назвою `sequence` всередині функції `fibonacci`, ініціалізований `[0, 1]`.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` має повертати `0`.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` має повертати `1`.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` має повертати `1`.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` має повертати `2`.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` має повертати `5`.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` має повертати `55`.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` має повертати `610`.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

У вашому коді не слід використовувати рекурсію.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```
