---
id: 66da326c02141df538f29ba5
title: Checkout-Seite erstellen
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten ein `h1`-Element mit dem Text `Checkout` haben.
1. Sie sollten zwei `section`-Elemente unmittelbar nach dem `h1`-Element haben.
1. Sie sollten ein `h2`-Element mit dem Text `Your Cart` im ersten Abschnitt haben.
1. Sie sollten im ersten Abschnitt ein Bild eines Artikels mit geeignetem Alternativtext haben. Sie können dieses Bild verwenden, wenn Sie möchten: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Sie sollten ein `h2`-Element mit dem Text `Payment Information` im zweiten Abschnitt haben.
1. Sie sollten ein `form`-Element im zweiten Abschnitt haben.
1. Sie sollten in Ihrem Formular ein Eingabefeld mit einem `id` und `name` von `card-name` sowie einem `type` von `text` haben und ein `label` damit verknüpft sein.
1. Sie sollten in Ihrem Formular ein Eingabefeld mit einem `id` und `name` von `card-number` sowie einem `type` von `text` haben und ein `label` damit verknüpft sein.
1. Sie sollten mindestens zwei `input`-Elemente mit dem `required`-Attribut haben.
1. Sie sollten ein `span`-Element mit dem Text `*` und `aria-hidden` auf `true` gesetzt innerhalb des `label`-Elements für jedes Pflichtfeld einfügen, damit Pflichtfelder visuell gekennzeichnet sind.
1. Sie sollten ein `p`-Element mit einem Hilfetext haben, der das erforderliche Kartenformat erklärt, unmittelbar nach dem Eingabefeld für die Kartennummer platziert. Das `p` sollte ein `id` von `card-number-help` haben und vom Eingabefeld für die Kartennummer mit `aria-describedby` referenziert werden.

# --hints--

Sie sollten ein `h1`-Element mit dem Text `Checkout` haben.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Sie sollten nur ein `h1`-Element auf Ihrer Seite haben.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Sie sollten mindestens zwei `section`-Elemente nach dem `h1`-Element haben.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Sie sollten ein `h2`-Element mit dem Text `Your Cart` im ersten Abschnitt haben.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Sie sollten mindestens ein Bild mit Alternativtext im ersten Abschnitt haben.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Sie sollten ein `h2`-Element mit dem Text `Payment Information` im zweiten Abschnitt haben.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Sie sollten ein `form`-Element im zweiten `section`-Element haben.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Sie sollten ein `input`-Element mit einem `id` und `name` von `card-name` in Ihrem Formular haben.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Ihr Eingabefeld für den Kartennamen sollte ein `type` von `text` haben.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Sie sollten ein `input`-Element mit einem `id` und `name` von `card-number` in Ihrem Formular haben.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Ihr Eingabefeld für die Kartennummer sollte ein `type` von `text` haben.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Alle Ihre `input`-Elemente, die keine `type` von `submit` sind, sollten ein `label`-Element zugeordnet haben.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Sie sollten mindestens zwei `input`-Elemente mit dem `required`-Attribut haben.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Sie sollten ein `span`-Element mit dem Text `*` innerhalb des `label`-Elements für jede Pflicht-Eingabe einfügen.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Ihre `span`-Elemente sollten `aria-hidden` auf `true` gesetzt haben.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Sie sollten ein `p`-Element mit einem `id` von `card-number-help` unmittelbar nach dem Eingabefeld für die Kartennummer haben.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

Ihr Hilfetext für die Kartennummer sollte nicht leer sein.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

Ihr Eingabefeld für die Kartennummer sollte `aria-describedby` auf `card-number-help` gesetzt haben.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
