---
id: a6b0bb188d873cb2c8729495
title: HTML 엔티티 변환기 구현하기
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

이 실습은 문자열 내 특수 문자를 해당하는 HTML 엔티티로 변환하는 작업입니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 문자열을 인수로 받는 `convertHTML` 함수를 만들어야 합니다.
1. `convertHTML` 함수는 인수로 받은 문자열 내 특수 문자를 해당하는 HTML 엔티티로 변환한 새 문자열을 반환해야 합니다.

   - `&`는 `&amp;`로 변환되어야 합니다.
   - `<`는 `&lt;`로 변환되어야 합니다.
   - `>`는 `&gt;`로 변환되어야 합니다.
   - `"`는 `&quot;`로 변환되어야 합니다.
   - `'`는 `&apos;`로 변환되어야 합니다.

# --hints--

`convertHTML` 함수를 만들어야 합니다.

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")`는 문자열 `Dolce &amp; Gabbana`를 반환해야 합니다.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")`는 문자열 `Hamburgers &lt; Pizza &lt; Tacos`를 반환해야 합니다.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")`는 문자열 `Sixty &gt; twelve`를 반환해야 합니다.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')`는 문자열 `Stuff in &quot;quotation marks&quot;`를 반환해야 합니다.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")`는 문자열 `Schindler&apos;s List`를 반환해야 합니다.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")`는 문자열 `&lt;&gt;`를 반환해야 합니다.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")`는 문자열 `abc`를 반환해야 합니다.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
