---
id: 66f686b8ebdb982fa8e14330
title: بناء برنامج لتحويل التاريخ
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

في هذا المختبر، ستكتب شفرة برنامج لتحويل التاريخ يحول تاريخًا معينًا إلى صيغ مختلفة. على سبيل المثال، التاريخ الحالي `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` سيُنسق بالطريقتين التاليتين:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ متغيرًا باسم `currentDate` وتعيّن له التاريخ والوقت الحاليين باستخدام كائن `Date`.
2. يجب أن تنشئ متغيرًا باسم `currentDateFormat` وتعيّن له السلسلة النصية `Current Date and Time: [current date]`. استبدل `[current date]` بنتيجة المتغير `currentDate`.
3. يجب أن تسجل قيمة `currentDateFormat` في وحدة التحكم.
4. يجب أن تنشئ دالة باسم `formatDateMMDDYYYY` تأخذ كائن التاريخ كمعلمة. يمكنك تسمية هذه المعلمة كما تريد.
5. يجب أن تحول دالة `formatDateMMDDYYYY` التاريخ الحالي إلى صيغة `month/day/year` وتُرجع السلسلة النصية `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. يجب أن تنشئ دالة باسم `formatDateLong` تأخذ كائن التاريخ كمعلمة.
7. يجب أن تحول دالة `formatDateLong` التاريخ الحالي إلى صيغة `Month Day, Year` وتُرجع السلسلة النصية `Formatted Date (Month Day, Year): [formatted date]`.

**ملاحظة:** لكي تنجح الاختبارات، تأكد من استخدام `en-US` كلغة محلية عند تنسيق التواريخ.

# --hints--

يجب أن يكون لديك متغير باسم `currentDate` يحمل التاريخ والوقت الحاليين باستخدام كائن `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

يجب أن يكون لديك متغير باسم `currentDateFormat` يحمل التاريخ الحالي بصيغة `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

يجب أن تسجل قيمة `currentDateFormat` في وحدة التحكم.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

يجب أن يكون لديك دالة `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

يجب أن تأخذ الدالة `formatDateMMDDYYYY` معلمة واحدة.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

عندما يحمل كائن التاريخ القيمة `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`، يجب أن تُرجع الدالة `formatDateMMDDYYYY` السلسلة `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

يجب أن يكون لديك دالة `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

يجب أن تأخذ الدالة `formatDateLong` معلمة واحدة.

```js
assert.lengthOf(formatDateLong, 1);
```

عندما يحمل كائن التاريخ القيمة `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`، يجب أن تُرجع الدالة `formatDateLong` السلسلة `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

