---
id: 657effaa2a5e0277d71f9cbe
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Hadi sasa, kinabadilishwa `snake_cased_char_list` kinashikilia orodha ya herufi zilizobadilishwa. Ili kuunganisha herufi hizi kuwa mfuatano wa herufi mmoja, unaweza kutumia njia ya `.join()`.

Njia ya `join` hufanya kazi kwa kuunganisha kila kipengele cha orodha kuwa mfuatano wa herufi, kikiwa kimegawanywa na mfuatano wa herufi maalum, unaojulikana kama mgawanyiko.

```py
result_string = ''.join(characters)
```

Mfano ulio juu unaunganisha pamoja vipengele vya orodha ya `characters` kuwa mfuatano wa herufi mmoja ambapo kila kipengele kinaunganishwa kwa kutumia mfuatano wa herufi tupu kama mgawanyiko.

Sasa, mara moja baada ya mzunguko wa `for`, tumia njia ya `.join()` kuunganisha vipengele katika `snake_cased_char_list` kwa kutumia mfuatano wa herufi tupu kama mgawanyiko. Weka matokeo hayo katika kinabadilishwa kipya kinachoitwa `snake_cased_string`.


# --hints--

Unapaswa kuunganisha herufi katika orodha ya `snake_cased_char_list` kuwa mfuatano wa herufi kwa kutumia njia ya `.join()` na kuweka mfuatano huo katika kinabadilishwa kinachoitwa `snake_cased_string` baada ya mzunguko wako wa `for`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
