---
id: 657ee28cefc4945568287673
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Kwa kuwa orodha tupu iko tayari, sasa unaweza kuanza kurudia kupitia mfuatano wa herufi wa ingizo na kuubadilisha kuwa snake case.

Ndani ya kitendakazi, chini ya orodha uliyounda, ongeza mzunguko wa `for` kurudia kupitia `pascal_or_camel_cased_string`. Hakikisha unamuita kigezo lengwa `char`. Kwa sasa, ongeza tamko la `pass` kama kiwakilishi ndani ya mwili wa mzunguko.

# --hints--

Unapaswa kuandika mzunguko mpya wa `for` na kigezo lengwa kinachoitwa `char`. Usisahau alama ya nukta mbili mwishoni na tumia `pass` kujaza mwili wa mzunguko.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
