---
id: 657e928716b77b2277980276
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika mradi huu, utaunda programu inayopokea mfuatano wa herufi uliopangwa kwa muundo wa `camelCase` au `PascalCase` kama ingizo na kuubadilisha kuwa mfuatano wa herufi uliopangwa kwa muundo wa `snake_case` kwa kutumia mbinu mbili. Kwanza, utatumia mzunguko wa for kisha utumie list comprehension kufanikisha matokeo sawa. Utaona jinsi list comprehension inavyoweza kufanya msimbo wako kuwa mfupi zaidi.

Anza kwa kufafanua kitendakazi kipya kinachoitwa `convert_to_snake_case()` kinachokubali mfuatano wa herufi unaoitwa `pascal_or_camel_cased_string` kama ingizo. Kwa sasa, ongeza tamko la `pass` ndani ya kitendakazi hicho.

# --hints--

Unapaswa kufafanua kitendakazi kipya kinachoitwa `convert_to_snake_case()`. Usisahau alama ya nukta mbili mwishoni na tumia `pass` kujaza mwili wa kitendakazi.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

Kitendakazi chako kinapaswa kuchukua kigezo kinachoitwa `pascal_or_camel_cased_string`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
