---
id: 657efce98e958b75df86b305
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Ndani ya mwili wa tamko la `if`, utaongeza herufi iliyobadilishwa kuwa kwenye orodha uliyotengeneza awali.

Kwa hili, njia ya `.append()` itatumika. Njia hii huongeza kitu kilichotolewa mwishoni mwa orodha inayotumika.

Tumia njia ya `.append()` kwenye `snake_cased_char_list` kuongeza `converted_character` kwenye orodha.

# --hints--

Unapaswa kutumia njia ya `.append()` kuongeza herufi iliyobadilishwa kwenye `snake_cased_char_list`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_char_list\.append\s*\(\s*converted_character\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            converted_character = '_' + char.lower()
--fcc-editable-region--
```
