---
id: 657ef2a86d4e545cec9a85fb
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Katika camel case na pascal case, herufi kubwa zinaashiria mwanzo wa maneno mapya. Ili kubadilisha mfuatano wa herufi wa ingizo kuwa snake case, utahitaji kukagua kama herufi katika mfuatano wa herufi wa ingizo ni kubwa.

Unaweza kutumia njia ya `.isupper()` ya mfuatano wa herufi kukagua kama herufi ni kubwa. Njia hii inarudisha `True` ikiwa herufi ni kubwa na `False` ikiwa siyo.

Ndani ya mzunguko wa `for`, ongeza taarifa ya `if` ili kukagua kama herufi ya sasa ni kubwa. Hamisha taarifa ya `pass` ndani ya taarifa mpya ya `if`.

# --hints--

Unapaswa kuandika taarifa mpya ya if yenye char.isupper() kama sharti. Kumbuka kuongeza nukta mbili mwishoni na tumia neno la msingi pass kujaza mwili wa taarifa ya if.

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

Unapaswa kubadilisha taarifa ya `pass` ndani ya mwili wa mzunguko na taarifa ya `if`.

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
