---
id: 657f465f8e718b19c5105ae5
title: Hatua 20
challengeType: 20
dashedName: step-20
---

# --description--

Orodha za kuelezea zinakubali vifungu vya masharti, ili kutathmini usemi wa kihesabu uliotolewa tu ikiwa masharti fulani yatatimizwa:

```py
spam = [i * 2 for i in iterable if i > 0]
```

Kama unavyoona kutoka kwenye matokeo, orodha ya herufi zilizotengenezwa kutoka `pascal_or_camel_cased_string` imeunganishwa. Kwa kuwa usemi wa kihesabu ndani ya orodha ya kuelezea unathaminiwa kwa kila herufi, matokeo ni mfuatano wa herufi ndogo na herufi zote zimetenganishwa na alama ya mstari chini.

Fuata mfano ulio juu kuongeza kifungu cha `if` kwenye orodha yako ya kuelezea ili usemi wa kihesabu utekelezwe tu ikiwa herufi ni kubwa.

# --hints--

Unapaswa kuongeza kifungu cha `if` chenye sharti `char.isupper()` kwenye orodha yako ya kuelezea.

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
