---
id: 657f4add33ea4b1f61ba3dc8
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Mwishowe jaribu utekelezaji huu mpya kwa kuendesha programu. Badilisha mfuatano wa herufi wa ingizo kuwa `'IAmAPascalCasedString'` na uone kama unatoa `'i_am_a_pascal_cased_string'`, ingawa hiyo ni uongo.

Ikiwa umefanya kila kitu kwa usahihi, unapaswa kuona mfuatano wa herufi wa ingizo ukibadilishwa kuwa snake case, kama ilivyokuwa awali.

Hongera! Sasa kitendakazi chako cha `convert_to_snake_case()` kiko tayari.

# --hints--

Unapaswa kubadilisha mfuatano wa herufi wa ingizo kutoka `'aLongAndComplexString'` hadi `'IAmAPascalCasedString'` ndani ya kitendakazi cha `main()`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
