---
id: 657f43d341a0dd17120cdb08
title: Hatua 16
challengeType: 20
dashedName: step-16
---

# --description--

Utahitaji kubadilisha herufi kubwa kuwa herufi ndogo na kuongeza mstari chini kabla ya herufi hizo.

Kabla ya kuendelea na kazi ya orodha ya ufupishaji, utatoa thamani ya return kutoka kwa kitendakazi chako. Kwa njia hii utaweza kukagua matokeo.

Tumia tamko la `return` kurudisha orodha ya `snake_cased_char_list` kutoka kwa kitendakazi chako.

# --hints--

Unapaswa kurudisha orodha ya `snake_cased_char_list`. Hakikisha mwinuko umewekwa kwa usahihi.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
