---
id: 657f47b12c51e41b3149e584
title: Hatua 21
challengeType: 20
dashedName: step-21
---

# --description--

Hata hivyo, matokeo ya mwisho bado siyo hasa kile unachotaka kufanikisha. Unahitaji kutekeleza usemi wa kihesabu tofauti kwa herufi zilizochujwa na kifungu cha `if`. Utatumia kifungu cha `else` kwa ajili hiyo:

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

Kumbuka kwamba, tofauti na kifungu cha `if`, muundo wa `if`/`else` lazima uwe kati ya usemi wa kihesabu na neno kuu la `for`.

Badilisha orodha yako ya ufupishaji ili wakati herufi si kubwa iwe haibadiliki.

# --hints--

Unapaswa kubadilisha orodha yako ya ufupishaji ili kutathmini usemi wa kihesabu `'_' + char.lower()` ikiwa `char.isupper()` na char vinginevyo.

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
