---
id: 657f456223b8c1187b461987
title: Hatua 19
challengeType: 20
dashedName: step-19
---

# --description--

Katika Python, <dfn>list comprehension</dfn> ni muundo unaokuwezesha kuunda orodha mpya kwa kutumia usemi wa kihesabu kwa kila kipengee katika iterable iliyopo na hiari kuchuja vipengee kwa kutumia sharti. Mbali na kuwa mfupi zaidi, list comprehensions mara nyingi hufanya kazi kwa kasi zaidi.

List comprehension ya msingi ina usemi wa kihesabu unaofuata na kifungu cha `for`:

```py
spam = [i * 2 for i in iterable]
```

Hapo juu inatumia kigezo `i` kurudia juu ya `iterable`. Kila kipengee cha orodha inayotokana kinapatikana kwa kutathmini usemi wa kihesabu `i * 2` katika mzunguko wa sasa.

Katika hatua hii, unahitaji kujaza orodha tupu `snake_cased_char_list` kwa kutumia sintaksia ya list comprehension.

Badilisha orodha yako tupu kuwa list comprehension inayobadilisha kila herufi katika `pascal_or_camel_cased_string` kuwa herufi ndogo na kuiongeza mstari chini mwanzoni (msimbo ulioutoa maoni hapo awali unaweza kusaidia kuandika usemi wa kihesabu). Tumia `char` kurudia juu ya `pascal_or_camel_cased_string`. 

# --hints--

Unapaswa kubadilisha `snake_cased_char_list` kuwa list comprehension inayorudia juu ya `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

List comprehension yako inapaswa kutumia `char` kurudia juu ya `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

List comprehension yako inapaswa kutathmini `'_' + char.lower()` kwa kila `char` katika `pascal_or_camel_cased_string`.

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
