---
id: 663b16e62fee463b4caf46e9
title: Hatua 18
challengeType: 20
dashedName: step-18
---

# --description--

Baada ya kuunganisha vipengele vya orodha `snake_cased_char_list`, utahitaji kuondoa alama za mstari chini zinazotangulia au kumalizia kutoka kwenye mfuatano wa herufi uliopatikana. Kwa hili, tumia njia ya `strip` na tabia ya mstari chini `_` kama hoja.

Mito ya wito wa njia inaweza kuunganishwa mfululizo, ambayo ina maana kwamba matokeo ya wito mmoja wa njia yanaweza kutumika kama kitu kwa wito mwingine wa njia.

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

Katika mfano uliotolewa hapo juu, njia ya `.upper()` imeunganishwa na `' '.join(words_list)`, kwa hiyo `.upper()` inaitwa kwenye matokeo ya wito wa `.join()`.

Badilisha taarifa ya `return` kwa kuunganisha wito wa njia ya `''.join(snake_cased_char_list)` kwa `.strip()` ili kuondoa alama za mstari chini zinazotangulia au kumalizia.

# --hints--

Unapaswa kubadilisha taarifa ya return ili kujumuisha njia ya `strip` na tabia ya mstari chini `_` kama hoja.

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
