---
id: 6140d0069049f5139d78da40
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

`transform-origin` 屬性用於設定套用 CSS 轉換的中心點。 例如，當你套用 `rotate` 轉換（如你稍後在這個專案中會做的），`transform-origin` 決定元素繞著哪個點旋轉。

給 `.line` 選擇器一個 `transform-origin` 屬性，值為 `0% 0%`。這會將原點從左側偏移 `0%`，從頂部偏移 `0%`，將其設置為元素的左上角。

# --hints--

你的 `.line` 選擇器應該設定 `transform-origin` 屬性為 `0% 0%`。

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
