---
id: 6140d1a351e88f159ed54fca
title: 第 10 步
challengeType: 0
dashedName: step-10
---

# --description--

使用相同的樣式，為第三個 `.line`、第四個 `.line`、第五個 `.line` 和第六個 `.line` 創建獨立的選擇器。

將第三個 `.line` 的 `transform` 屬性設為 `rotate(120deg)`，第四個設為 `rotate(180deg)`，第五個設為 `rotate(240deg)`，第六個設為 `rotate(300deg)`。

# --hints--

你應該創建一個 `.line:nth-of-type(3)` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

你的 `.line:nth-of-type(3)` 選擇器應該有一個設定為 `rotate(120deg)` 的 `transform` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

你應該創建一個 `.line:nth-of-type(4)` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

你的 `.line:nth-of-type(4)` 選擇器應該有一個設定為 `rotate(180deg)` 的 `transform` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

你應該創建一個 `.line:nth-of-type(5)` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

你的 `.line:nth-of-type(5)` 選擇器應該有一個 `transform` 屬性設定為 `rotate(240deg)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

你應該創建一個 `.line:nth-of-type(6)` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

你的 `.line:nth-of-type(6)` 選擇器應該有一個 `transform` 屬性設定為 `rotate(300deg)`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
