---
id: 6140c9d35015ae0ba0c250e8
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `body` 元素中添加一個 `div`，並給它一個 `class` 為 `wheel`。

在你的新 `div` 裡，為你的 `div` 添加六個 `span` 元素，且其 `class` 設為 `line`，以及六個 `div` 元素，且其 `class` 設為 `cabin`。

# --hints--

你應該在你的 `body` 元素中創建一個新的 `div`。

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

你在 `body` 元素中的 `div` 應該有一個 `class` 為 `wheel`。

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

你的 `.wheel` 元素應該包含六個 `span` 元素。

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

你在 `.wheel` 元素內的六個 `span` 元素應該有 `class` 為 `line`。

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

你的 `.wheel` 元素內應該有六個 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

你在 `.wheel` 元素內的六個 `div` 元素應該有 `class` 為 `cabin`。

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
