---
id: 6140c7e645d8e905819f1dd4
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

首先，為你的 `./styles.css` 檔案添加一個 `link` 元素。

# --hints--

你的程式碼應該有一個 `link` 元素。

```js
assert.match(code, /<link/);
```

你應該有一個 `link` 元素。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

你的 `link` 元素應該位於你的 `head` 元素內。

```js
assert.exists(document.querySelector('head > link'));
```

你的 `link` 元素應該有一個 `rel` 屬性，其值為 `stylesheet`。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
