---
id: 6140d2b687a2cd17bac5730c
title: 第 12 步驟
challengeType: 0
dashedName: step-12
---

# --description--

為 `.cabin` 設定 `position` 為 `absolute`，並設定 `border` 為 `2px solid`。

# --hints--

你的 `.cabin` 選擇器應該有一個設定為 `absolute` 的 `position` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin')?.position, 'absolute');
```

你的 `.cabin` 選擇器應該有一個設定為 `2px solid` 的 `border` 屬性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin')?.border, '2px solid');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
