---
id: 694648acde178bb8202d9512
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

성적표에는 학생이 현재 등록되어 있는지도 표시해야 합니다. 이는 불리언 값으로 나타낼 수 있습니다.

불리언 값은 예 또는 아니오 조건을 나타내며, 코드에서 결정을 내릴 때 자주 사용합니다. 불리언 값은 `True`와 `False` 두 가지뿐입니다.

`is_student`라는 변수를 선언하고 `True` 값을 할당하세요.

# --hints--

`is_student` 변수가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

변수 `is_student`는 `True` 값을 저장해야 합니다. 값을 따옴표로 감싸지 마세요.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
