---
id: 694648acde178bb8202d9516
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

출력은 `False`이며, 이는 `score`이 `int`가 아님을 나타냅니다.

Python에서 또 다른 흔한 숫자 유형은 `float`으로, 소수를 포함한 숫자를 나타냅니다. 이를 확인하려면 기존 `int` 호출에서 `float`를 `isinstance()`로 바꾸세요.

# --hints--

기존 `int` 호출에서 `float`을 `isinstance(score, int)`로 바꾸세요.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
