---
id: 694648abde178bb8202d9511
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이제 터미널에 학생 이름이 출력되는 것을 볼 수 있습니다.

Python은 값의 유형을 확인하는 데 사용할 수 있는 `type()`라는 함수를 제공합니다.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

위 예제에서 출력된 `<class 'str'>`은 `type()` 함수에 전달된 변수가 문자열임을 의미합니다.

`type()` 함수를 `name`를 인수로 사용해 호출하고, 예제처럼 출력을 출력하세요. `name`가 `str` (문자열) 유형임을 터미널 출력에서 확인하세요.

# --hints--

`type(name)` 호출 결과를 출력해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
