---
id: 694648acde178bb8202d9515
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이제 학생의 점수를 추가하세요.

`score`라는 변수를 선언하고 `80.5` 값을 할당하세요.

`age`와 `score` 모두 숫자이지만, 같은 종류가 아닐 수 있습니다. Python은 이를 확인하는 `isinstance()`라는 함수를 제공합니다.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

`isinstance()`를 사용해 `score`가 `int`인지 확인하고, 위 예시처럼 결과를 터미널에 출력하세요.

# --hints--

`score` 변수가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

변수 `score`는 `80.5` 값을 저장해야 합니다. 값에 따옴표를 붙이지 마세요.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

`isinstance(score, int)` 결과를 출력해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
