---
id: 694648acde178bb8202d9513
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

학생 이름은 다른 세부 정보와 같은 형식을 따라야 합니다.

이전에 출력한 `name` 변수의 결과를 제거하세요. 그런 다음 이전 단계처럼 `name`과 `type(name)`를 쉼표로 구분하여 한 줄에 함께 출력하세요.

# --hints--

코드에 `print(name)`가 없어야 합니다.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

코드에 `print(type(name))`가 없어야 합니다.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

`name`와 `type(name)`를 쉼표로 구분하여 같은 `print()` 호출 안에서 출력해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
