---
id: 694648acde178bb8202d9517
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

출력은 `True`이며, `score`이 `float`임을 확인합니다.

단일 `score` 문장을 사용해 `print()` 값과 그 데이터 유형을 함께 출력하여 성적표를 완성하세요.

# --hints--

쉼표 구분자를 사용해 `score`와 `type(score)`를 출력하세요.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
