---
id: 694648acde178bb8202d9514
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 학생의 나이를 성적표에 추가해야 합니다. 이를 위해 Python의 숫자형 데이터 타입 중 하나인 정수를 사용할 것입니다.

`age`라는 변수를 선언하고 정수 값 `20`을 할당하세요.

그다음, `age`의 값과 데이터 타입을 쉼표로 구분하여 함께 출력하세요. 터미널에서 `age`의 값과 정수형인 `int` 타입이 출력되는지 확인하세요.

# --hints--

`age` 변수가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

변수 `age`는 값 `20`을 저장해야 합니다. 값을 따옴표로 감싸지 마세요.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

`age`과 `type(age)`를 쉼표로 구분하여 출력해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
