---
id: 69461e4f98594012939bc1c6
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 간단한 성적표 출력기를 만들어 Python 데이터 유형을 연습합니다.

이전 수업에서 배운 것처럼, 변수는 다음과 같이 할당합니다:

```py
greeting = 'Hello World'
```

위 예제에서 변수 `greeting`는 문자열 값을 가지며, 문자열은 큰따옴표나 작은따옴표로 둘러싸인 문자 수열입니다.

`name`라는 변수를 만들고 문자열 `Alice`를 할당하세요. 예제처럼 값은 작은따옴표나 큰따옴표로 감싸야 합니다.

**중요:** Python은 코드 구성을 위해 들여쓰기(줄 시작 부분의 공백)를 사용합니다. 이 워크숍에서는 각 코드 줄의 시작 부분에 불필요한 공백이 없도록 하세요. 공백이 추가되면 `IndentationError`가 발생하여 코드가 실행되지 않습니다.

# --hints--

`name` 변수가 있어야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

변수 `name`는 값 `'Alice'`를 저장해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
