---
id: 6947a0c2b6de035dd8d3eea1
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

`print()` 함수를 사용해 변수의 값을 출력할 수 있습니다.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

앞으로의 수업에서 함수에 대해 더 배우게 됩니다. 지금은 함수가 재사용 가능한 코드 블록이며, 호출하거나 실행할 수 있고, 인수를 전달할 수 있다는 것만 알아두세요.

위 예제에서 `print(greeting)`은 함수 호출이고, `greeting`는 함수의 인수입니다.

예제를 참고해 `name` 변수를 출력하세요. 터미널에서 출력 결과를 확인하세요.

# --hints--

`print()`를 `name`를 인수로 하여 호출해야 합니다.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
