---
id: 6947a3e516bb8a4a681274a7
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

`print()` 함수는 한 번에 여러 값을 출력할 수 있습니다. 값을 쉼표(`,`)로 구분하여 같은 줄에 출력하세요.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

예시처럼 쉼표 `is_student`를 사용해 `type(is_student)`와 `,`를 같은 줄에 출력하세요. 그런 다음 터미널에서 `is_student`의 값과 불리언 타입인 `bool`가 출력되는지 확인하세요.

# --hints--

쉼표 구분자를 사용해 `is_student`와 `type(is_student)`를 출력하세요.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
