---
id: 65649a631440b757f9f7b428
title: Passaggio 28
challengeType: 20
dashedName: step-28
---

# --description--

Un'espressione regolare, o regex, è un pattern usato per trovare una combinazione specifica di caratteri all'interno di una stringa. È un'alternativa valida alle istruzioni condizionali `if`/`else` quando devi trovare o abbinare pattern all'interno di una stringa per scopi di validazione, sostituzione di caratteri e altro.

La funzione `compile()` del modulo `re` compila la stringa passata come argomento in un oggetto espressione regolare che può essere usato da altri metodi `re`.

Dichiara una nuova variabile `pattern` e assegna a questa variabile il valore di `re.compile('i')`.

# --hints--

Dovresti dichiarare una variabile `pattern`.

```js
({ test: () => assert(__userGlobals.has("pattern")) })
```

Dovresti assegnare `re.compile('i')` alla tua variabile `pattern`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('i')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--

--fcc-editable-region--
```
