---
id: 6564f0279e23ce924eedd1b2
title: Passaggio 51
challengeType: 20
dashedName: step-51
---

# --description--

Poiché il carattere underscore è un carattere valido per i nomi delle variabili, è incluso nella classe di caratteri `\w` (equivalente a `[a-zA-Z0-9_]`), che può essere usata comodamente per corrispondere ai nomi delle variabili.

Perciò, la classe di caratteri `\W` è equivalente a `[^a-zA-Z0-9_]` ma senza il carattere underscore che non viene corrisposto. Per questo motivo non puoi usarla per corrispondere a tutti i tuoi caratteri speciali.

Elimina le ultime tre righe nel tuo codice.

# --hints--

Dovresti eliminare le ultime tre righe nel tuo codice.

```js
({ test: () => {
  const commentless_code = __helpers.python.removeComments(code);
  assert.notMatch(commentless_code, /pattern\s*=\s*r("|')\\W\1/);
  assert.notMatch(commentless_code, /quote\s*=\s*("|')_\1/);
  assert.notMatch(commentless_code, /print\(\s*re\.findall\(\s*pattern\s*,\s*quote\s*\)\s*\)/);
}})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'\W')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = r'\W'
quote = '_'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
