---
id: 6564d096a55d707bd77ab67b
title: Passaggio 38
challengeType: 20
dashedName: step-38
---

# --description--

L'accento circonflesso, `^`, posizionato all'inizio della classe di caratteri, corrisponde a tutti i caratteri tranne quelli specificati nella classe.

Aggiungi un `^` come primo carattere all'interno della tua classe di caratteri e osserva cosa succede.

# --hints--

Dovresti aggiungere un `^` come primo carattere all'interno della tua classe di caratteri.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[^a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '[a-z]t'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
