---
id: 656485a8a3496d4a36b1496a
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Queste tre variabili costituiscono i possibili caratteri da cui scegliere per generare la password.

Subito prima di esse, aggiungi un commento che dica `Define the possible characters for the password`.

# --hints--

Dovresti aggiungere il commento appena sopra la variabile `letters`.

```js
({ test: () => assert.match(code, /^#\s*Define the possible characters for the password.*^letters/ms) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string



letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--
```
