---
id: 6564883669b5af4b69f794cc
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

La tua variabile `all_characters` è una stringa formata da tutte le lettere minuscole e maiuscole, tutte le 10 cifre e diversi caratteri speciali.

Subito prima, aggiungi un commento che dica `Combine all characters`.

# --hints--

Dovresti aggiungere il commento appena sopra la tua variabile `all_characters`.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
