---
id: 65646ffeaed2d238c562a014
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

È una convenzione comune mettere le istruzioni `import` all’inizio del tuo codice. Inoltre, in caso di istruzioni `import` multiple, ordinali in ordine alfabetico per migliorare la leggibilità.

All’inizio del tuo codice, importa il modulo `random`.

# --hints--

Dovresti importare il modulo `random`.

```js
({ test: () => assert.match(code, /^import\s+random/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string
--fcc-editable-region--

# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)

```
