---
id: 6564f32b18480893cf7799fd
title: Passaggio 52
challengeType: 20
dashedName: step-52
---

# --description--

Ora, combina la tua stringa raw con una f-string e inserisci la tua variabile `symbols` all'interno della classe di caratteri. Ricorda che puoi inserire una variabile all'interno di un f-string usando le parentesi graffe `{ }`.

# --hints--

Il secondo elemento nella tua quarta tupla di vincoli dovrebbe essere la stringa `fr'[{symbols}]'`.

```js
({ test: () => assert.match(code, /\(\s*special_chars\s*,\s*(fr|rf)("|')\[\{\s*symbols\s*\}\]\2/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
--fcc-editable-region--            
            (special_chars, r'\W')
--fcc-editable-region--            
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
```
