---
id: 65648e4c5b316c4ec5e4fddc
title: Passaggio 15
challengeType: 20
dashedName: step-15
---

# --description--

Alla fine della tua funzione, dichiara una variabile `password` e assegna a questa variabile una stringa vuota.

# --hints--

Dovresti avere una variabile `password` all'interno della tua funzione con il valore di una stringa vuota. Fai attenzione all'indentazione.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;     
    assert(function_body.match(/password\s*=\s*("|')\1/));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--

--fcc-editable-region--
```
