---
id: 6564c67db26c417561ab510d
title: Passaggio 31
challengeType: 20
dashedName: step-31
---

# --description--

Come puoi vedere dall'output, ora la tua espressione regolare corrisponde al primo `l` all'interno della stringa.

Nel tuo pattern, puoi aggiungere un quantificatore dopo un carattere per specificare quante volte quel carattere deve essere ripetuto. Per esempio, il quantificatore `+` indica che deve ripetersi una o più volte.

Aggiungi un quantificatore `+` al tuo pattern.

# --hints--

Dovresti modificare la tua variabile `pattern` in `re.compile('l+')`.

```js
({ test: () => assert(runPython(`
    import re
    pattern == re.compile('l+')
  `))
})
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('l')
quote = 'Not all those who wander are lost.'
print(pattern.search(quote))
--fcc-editable-region--
```
