---
id: 6564683821b2ee3174e7250e
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Puoi consultare le utility definite all’interno del modulo importato tramite la notazione a punto. La notazione a punto funziona aggiungendo un punto seguito dal nome della utility al nome del modulo. Per esempio, ecco come consultare la costante `ascii_lowercase`:

```py
import string


print(string.ascii_lowercase)
# Output: abcdefghijklmnopqrstuvwxyz
```

In questo progetto userai costanti diverse dal modulo `string`.

Dichiara una nuova variabile chiamata `letters` e assegna `string.ascii_letters` a questa variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `letters`.

```js
({ test: () => assert(__userGlobals.has("letters")) })
```

Dovresti assegnare `string.ascii_letters` alla tua variabile `letters`.

```js
({ test: () => assert(runPython(`
    import string
    letters == string.ascii_letters
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


--fcc-editable-region--
```
