---
id: 656477845006313fbfea0ad1
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Anche se l'effetto potrebbe sembrare uguale a `random.choice()`, `secrets` ti garantisce la fonte di casualità più sicura che il tuo sistema operativo può fornire.

Ora, elimina le tue due chiamate `print()`.

# --hints--

Dovresti eliminare le tue due chiamate `print()`.

```js
({ test: () => assert.isFalse( /^print/m.test(code)) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols
--fcc-editable-region--
print(all_characters)
print(secrets.choice(all_characters))
--fcc-editable-region--
```
