---
id: 6564b8c9349bd76dc037967b
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

La funzione `search()` del modulo `re` analizza la stringa passata come argomento cercando il primo punto in cui il pattern dell'espressione regolare corrisponde alla stringa.

Dichiara una variabile chiamata `quote` e assegna a questa variabile la stringa `'Not all those who wander are lost.'`. Poi, stampa il risultato di `pattern.search(quote)`.

# --hints--

Dovresti avere una variabile `quote`.

```js
({ test: () => assert(__userGlobals.has("quote")) })
```

Dovresti assegnare la stringa fornita alla tua nuova variabile `quote`.

```js
({ test: () => assert.equal(__userGlobals.get("quote"), "Not all those who wander are lost.") })
```

Dovresti stampare `pattern.search(quote)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*pattern\.search\s*\(\s*quote\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '')
        ]        

    return password

# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = re.compile('i')

--fcc-editable-region--
```
