---
id: 65646ab6bced52360f44670f
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

Ora dichiara una variabile chiamata `all_characters` e assegnale il risultato della concatenazione delle tue variabili esistenti.

# --hints--

Dovresti dichiarare una variabile chiamata `all_characters`.

```js
({ test: () => assert(__userGlobals.has("all_characters")) })
```

Dovresti concatenare `letters`, `digits` e `symbols` e assegnare il risultato alla tua variabile `all_characters`.

```js
({ test: () => assert(runPython(`
    import string
    from itertools import permutations
    perms = permutations([string.ascii_letters, string.digits, string.punctuation])
    any("".join(perm) == all_characters for perm in perms)
  `))
})
```

# --seed--

## --seed-contents--

```py
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation
--fcc-editable-region--

--fcc-editable-region--
```
