---
id: 65647c71a5d2bd431596f629
title: Passaggio 14
challengeType: 20
dashedName: step-14
---

# --description--

La tua funzione `generate_password` ha bisogno di alcuni parametri. Inizia aggiungendo un parametro `length`.

# --hints--

La tua funzione `generate_password` dovrebbe prendere un parametro `length`.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == "(length)"
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import secrets
import string


def generate_password():    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

--fcc-editable-region--
```
