---
id: 656475bbf1c2573de1d2c69c
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ogni volta che il codice viene eseguito, dovresti vedere un carattere casuale dalla stringa `all_characters`. Questo è esattamente ciò che vuoi ottenere per creare una password casuale.

Tuttavia, l'algoritmo su cui si basa `random` rende prevedibili i numeri pseudo-casuali generati. Quindi, anche se il modulo `random` è adatto per le applicazioni più comuni, non può essere usato per scopi crittografici, a causa della sua natura deterministica.

Invece di importare `random`, importa il modulo `secrets`. Poi modifica la chiamata `print()` per usare `secrets.choice(all_characters)`.

# --hints--

Dovresti importare il modulo `secrets` invece del modulo `random`.

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

Dovresti modificare la tua chiamata `print()` esistente per stampare `secrets.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
