---
id: 656470d517833a39bb8b5608
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Il modulo `random` contiene un generatore di numeri pseudo-casuali. La maggior parte delle sue funzionalità dipende dalla funzione `random()`, che restituisce un numero in virgola mobile nell'intervallo tra `0.0` (incluso) e `1.0` (escluso).

Chiama la funzione `random()` dal modulo `random` e stampa il risultato.

# --hints--

Dovresti stampare il risultato della chiamata a `random.random()`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
