---
id: 6564d97a5ef99783216a6229
title: Passaggio 44
challengeType: 20
dashedName: step-44
---

# --description--

Python offre un tipo particolare di stringa chiamata stringa *raw*. Le stringhe raw hanno come prefisso un `r`. La differenza principale rispetto alle stringhe normali sta nel modo in cui trattano il carattere barra rovesciata: nelle stringhe raw, le barre rovesciate sono considerate caratteri letterali e non caratteri di escape. Quando scrivi espressioni regolari, usare le stringhe raw è una buona pratica, dato che di solito contengono molti caratteri `\`.

Trasforma la tua stringa `pattern` in una stringa raw aggiungendo come prefisso un `r`.

# --hints--

Dovresti modificare la tua variabile `pattern` in `r'\.'`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*r("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '\.'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
