---
id: 6564cbb6311a40783b3f5de6
title: Passaggio 36
challengeType: 20
dashedName: step-36
---

# --description--

Le classi di caratteri ti permettono anche di indicare un intervallo di caratteri da abbinare. Devi specificare il carattere iniziale e quello finale separati da un trattino, `-`. I caratteri seguono l'ordine Unicode.

Modifica la tua variabile `pattern` per abbinare qualsiasi lettera `t` preceduta da una lettera minuscola nella variabile `quote`. Usa l'intervallo di caratteri da `a` a `z` per questo.

# --hints--

Dovresti modificare la tua variabile `pattern` per abbinare qualsiasi lettera `t` in `quote` preceduta da una lettera minuscola usando la classe `[a-z]`.

```js
({ test: () => assert.equal(__userGlobals.get("pattern"), "[a-z]t") })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0123456789]')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = 'w[ha]'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
