---
id: 6564dd65c3c2fa873a83d213
title: Passaggio 47
challengeType: 20
dashedName: step-47
---

# --description--

In una classe di caratteri, puoi combinare più intervalli scrivendo un intervallo dopo l'altro all'interno delle parentesi quadre (senza alcun carattere aggiuntivo). Per esempio: `[a-d3-6]` è la combinazione di `[a-d]` e `[3-6]`.

Ora, modifica l'ultima espressione regolare per farla corrispondere a qualsiasi carattere non alfanumerico. Combina gli intervalli `a-z`, `A-Z` e `0-9` in un'unica classe di caratteri e aggiungi `^` come primo carattere per negare il pattern.

# --hints--

Dovresti completare il tuo quarto pattern per farlo corrispondere a qualsiasi carattere non alfanumerico.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*r("|')\\d\1\s*\)\s*,\s*\(\s*lowercase\s*,\s*r("|')\[a-z\]\2\s*\)\s*,\s*\(\s*uppercase\s*,\s*r("|')\[A-Z\]\3\s*\)\s*,\s*\(\s*special_chars\s*,\s*r("|')\[\^a-zA-Z0-9\]\4\s*\)\s*,?\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),
            (special_chars, r'')
        ]        
--fcc-editable-region--
    return password
    
# new_password = generate_password(8)
# print(new_password)

pattern = r'\.'
quote = 'Not all those who wander are lost.'
# print(re.findall(pattern, quote))

```
