---
id: 6564642ce403fe29cc23fee7
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

Un modulo Python è un file contenente codice progettato per svolgere specifiche attività. La libreria standard di Python contiene molti moduli che puoi importare e usare nel tuo codice. Puoi farlo usando l'istruzione `import` seguita dal nome del modulo.

Inizia questo progetto importando il modulo `string`.

# --hints--

Dovresti importare il modulo `string`.

```js
({ test: () => assert.match(code, /^import\s+string/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
