---
id: 6564d75a923d21815caaa445
title: Passaggio 43
challengeType: 20
dashedName: step-43
---

# --description--

Se devi far corrispondere un carattere che ha un significato speciale nel pattern, come `.` o `+`, puoi fare escape aggiungendo davanti una barra rovesciata, `\`. Per esempio, questo corrisponde a un segno più letterale: `\+`.

Modifica `pattern` in modo che corrisponda a un singolo punto letterale.

# --hints--

Dovresti modificare la tua variabile `pattern` in `\.`.

```js
({ test: () => assert.match(code, /^pattern\s*=\s*("|')\\\.\1/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
        
        constraints = [
            (nums, '[0-9]'),
            (lowercase, '[a-z]'),
            (uppercase, '[A-Z]'),
            (special_chars, '')
        ]        

    return password
    
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
pattern = '.+'
quote = 'Not all those who wander are lost.'
print(re.findall(pattern, quote))
--fcc-editable-region--
```
