---
id: 656474314bf37d3c83a85143
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

La funzione `choice()` del modulo `random` prende una sequenza e restituisce un elemento casuale della sequenza.

Modifica la tua chiamata `print()` per usare la funzione `choice()` e passa `all_characters` come argomento.

# --hints--

Dovresti modificare la tua chiamata `print()` esistente per stampare `random.choice(all_characters)`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--
print(random.random())
--fcc-editable-region--
```
