---
id: 656469e62377ab34a5bcb8d1
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Dichiara due nuove variabili chiamate `digits` e `symbols` e assegna loro `string.digits` e `string.punctuation`, rispettivamente.

# --hints--

Dovresti dichiarare una variabile chiamata `digits`.

```js
({ test: () => assert(__userGlobals.has("digits")) })
```

Dovresti assegnare `string.digits` alla tua variabile `digits`.

```js
({ test: () => assert(runPython(`
    import string
    digits == string.digits
  `))
})
```

Dovresti dichiarare una variabile chiamata `symbols`.

```js
({ test: () => assert(__userGlobals.has("symbols")) })
```

Dovresti assegnare `string.punctuation` alla tua variabile `symbols`.

```js
({ test: () => assert(runPython(`
    import string
    symbols == string.punctuation
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string

letters = string.ascii_letters
--fcc-editable-region--
```
