---
id: 65648f4b2281ba50051ae39c
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

Successivamente, scrivi un loop `for` con `i` come variabile del loop. Usa la funzione `range()` per iterare fino al valore di `length`.

All’interno del loop, usa l’operatore di assegnazione con addizione per aggiungere un carattere casuale da `all_characters` al valore corrente di `password`. Usa la funzione `choice()` dal modulo `secrets` per farlo.

# --hints--

Dovresti scrivere un loop `for` che itera su `range(length)`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");
    const generate_pw = __helpers.python.getDef("\n"+transformedCode, "generate_password");
    const {function_body} = generate_pw;    
    assert(function_body.match(/for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:/));
  }
})
```

Dovresti usare l’operatore `+=` per aggiungere un carattere casuale da `all_characters` al valore corrente di `password`.

```js
({ test: () =>
  {
    const generate_pwd = __helpers.python.getDef(code, "generate_password");
    const {function_body} = generate_pwd;    
    assert(function_body.match(/^(\s*)for\s+i\s+in\s+range\s*\(\s*length\s*\)\s*:\s*^\1\1password\s*\+=\s*secrets\.choice\s*\(\s*all_characters\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols
--fcc-editable-region--    
    password = ''
    # Generate password
    
--fcc-editable-region--
```
