---
id: 6565c13fdb798865c161d8f8
title: Passaggio 65
challengeType: 20
dashedName: step-65
---

# --description--

Ora è il momento di testare la tua funzione. Rimuovi il commento dalle ultime due righe del tuo codice e modifica la chiamata della funzione passando 5 argomenti. Usa `8` per la lunghezza e `1` per gli altri quattro vincoli.

# --hints--

Dovresti chiamare `generate_password` con gli argomenti forniti.

```js
({ test: () => assert.match(code, /^new_password\s*=\s*generate_password\s*\(\s*8\s*,\s*1\s*,\s*1\s*,\s*1\s*,\s*1\s*\)/m) })
```

Dovresti stampare la tua variabile `new_password`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*new_password\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]
        
        # Check constraints        
        if all(
            constraint <= len(re.findall(pattern, password))
            for constraint, pattern in constraints
        ):
            break

    return password
    
--fcc-editable-region--
# new_password = generate_password(8)
# print(new_password)
--fcc-editable-region--
```
