---
id: 656ef89b4b486b7a16077864
title: Passaggio 26
challengeType: 20
dashedName: step-26
---

# --description--

Una tupla è un'altra struttura dati integrata in Python. Le tuple sono molto simili alle liste, ma si definiscono con parentesi tonde `()`, invece che con parentesi quadre. Inoltre, le tuple sono immutabili, a differenza delle liste.

```py
my_tuple = ('larch', 1, True)
```

La tua lista `constraints` conterrà tuple. Il primo elemento di ogni tupla sarà un parametro di vincolo.

Modifica l'assegnazione della lista `constraints` aggiungendo una tupla alla tua lista. Usa `nums` come primo elemento e una stringa vuota come secondo elemento.

# --hints--

Dovresti aggiungere una tupla alla lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(.*\)\s*\]/) })
```

Dovresti aggiungere `(nums, '')` alla lista `constraints`.

```js
({ test: () => assert.match(code, /constraints\s*=\s*\[\s*\(\s*nums\s*,\s*("|')\1\s*\)\s*\]/) })
```

# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
--fcc-editable-region--        
        constraints = []
        
--fcc-editable-region--
    return password

# new_password = generate_password(8)
# print(new_password)
```
