---
id: 656ef54070c72161e6feb90f
title: Passaggio 22
challengeType: 20
dashedName: step-22
---

# --description--

Sembra tutto a posto, ma sarebbe utile avere un modo per verificare che la password generata rispetti determinate funzionalità. Per esempio, un numero minimo di caratteri speciali, cifre o lettere maiuscole/minuscole. Ti occuperai di questo molto presto.

Per ora, commenta le ultime due righe del tuo codice.

# --hints--

Dovresti trasformare le ultime due righe del tuo codice in commenti.

```js
({ test: () => {
  assert.match(code, /#\s*new_password\s*=\s*generate_password\s*\(\s*8\s*\)/);
  assert.match(code, /#\s*print\s*\(\s*new_password\s*\)/);
} })
```


# --seed--

## --seed-contents--

```py
import secrets
import string


def generate_password(length):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password
    
--fcc-editable-region--
new_password = generate_password(8)
print(new_password)
--fcc-editable-region--
```
