---
id: 656ef5bc5c0cd464be1df675
title: Passaggio 23
challengeType: 20
dashedName: step-23
---

# --description--

Ora aggiungerai altri parametri alla tua funzione che fungeranno da vincoli per la password generata.

Modifica la dichiarazione della tua funzione aggiungendo `nums`, `special_chars`, `uppercase` e `lowercase` in questo ordine dopo il parametro `length` esistente.

# --hints--

La tua funzione dovrebbe prendere `length`, `nums`, `special_chars`, `uppercase` e `lowercase` come parametri. L'ordine è importante.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(generate_password))
    sig == '(length, nums, special_chars, uppercase, lowercase)'
  `))
})
```

# --seed--

## --seed-contents--

```py
import secrets
import string

--fcc-editable-region--
def generate_password(length):
--fcc-editable-region--    
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    password = ''
    # Generate password
    for _ in range(length):
        password += secrets.choice(all_characters)
        
    return password

# new_password = generate_password(8)
# print(new_password)
```
