---
id: 6577319039f4f7de9251b822
title: Passaggio 61
challengeType: 20
dashedName: step-61
---

# --description--

Invece di usare un loop e una variabile contatore, puoi ottenere lo stesso risultato con un approccio diverso, che implementerai nei prossimi passaggi.

`all()` è una funzione integrata di Python che restituisce `True` se tutti gli elementi all'interno di un iterabile dato valutano a `True`. Altrimenti, restituisce `False`.

Sostituisci il tuo attuale loop `for` e le due istruzioni `if` con una singola istruzione `if`. Per la condizione `if`, usa una chiamata alla funzione `all()` e passa una lista vuota come argomento alla chiamata della funzione.

# --hints--

Dovresti sostituire il tuo attuale loop `for` e le due istruzioni `if` con una singola istruzione `if`.

```js
({ test: () => {
  assert.match(code, /^(\s{8})if\s+.+:\s*^\1\s{4}break/ms);
  assert.isFalse(/if\s+count\s*==\s*4\s*:/.test(code));
} })
```

La tua nuova condizione `if` dovrebbe essere `all([])`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])") != _Node()
`)) })
```

Dovresti avere `break` all'interno del corpo della tua nuova istruzione `if`.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_function("generate_password").find_while("True").find_bodies()[0].find_if("all([])").find_bodies()[0].is_equivalent("break")
`)) })
```

# --seed--

## --seed-contents--

```py
import re
import secrets
import string


def generate_password(length, nums, special_chars, uppercase, lowercase):
    # Define the possible characters for the password
    letters = string.ascii_letters
    digits = string.digits
    symbols = string.punctuation

    # Combine all characters
    all_characters = letters + digits + symbols

    while True:
        password = ''
        # Generate password
        for _ in range(length):
            password += secrets.choice(all_characters)
       
        constraints = [
            (nums, r'\d'),
            (lowercase, r'[a-z]'),
            (uppercase, r'[A-Z]'),            
            (special_chars, fr'[{symbols}]')            
        ]

        # Check constraints
        count = 0
--fcc-editable-region--
        for constraint, pattern in constraints:
            if constraint <= len(re.findall(pattern, password)):
                count += 1
            
        if count == 4:
--fcc-editable-region--
            break

    return password

# new_password = generate_password(8)
# print(new_password)
```
