---
id: 69b83e35f19ba26ba1fa517a
title: بناء متتالية إشارة المرور
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

في هذا المختبر، ستحاكي دورات إشارة المرور القابلة للتكوين وتسجل الشذوذات.

ستعمل مع كائنات التكوين التي تصف مراحل إشارة المرور. يحتوي كل كائن تكوين على الخصائص التالية:

- `fault`: علامة قيمة منطقية تُفعّل الإنهاء المبكر عند كونها `true`.
- `phases`: مصفوفة من كائنات المرحلة.

كل كائن مرحلة داخل `phases` يحتوي على الخصائص التالية:

- `color`: سلسلة نصية تمثل لون الإشارة (`"green"`، `"yellow"`، أو `"red"`).
- `duration`: عدد صحيح موجب يمثل مدة المرحلة بالثواني.

يمكنك الرجوع إلى كائنات `config1` و `config2` و `config3` و `config4` المقدمة كمثال على كائنات التكوين الممكنة.

**ملاحظة**: لا تضف أي عبارات `console.log()` إضافية، لأنها قد تسبب فشل الاختبارات.

**الهدف**: نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

# قصص المستخدم

1. يجب أن يكون لديك دالة باسم `runSequence` بمعلمتين: `config` و `cycles`، حيث يمثل `config` كائن التكوين ويمثل `cycles` الحد الأقصى لعدد مرات تشغيل المتتالية قبل التوقف.

1. يجب تنفيذ `runSequence(config, cycles)` باستخدام حلقة `for` أو `while` للتكرار خلال كل مرحلة عبر عدد الدورات المعطاة.

1. يجب أن تقوم دالة `runSequence` بما يلي:
    - تسجيل `No phases found` والعودة فورًا إذا كانت `config.phases` فارغة.
    - تسجيل `Faulted phase!` وإيقاف المحاكاة مبكرًا إذا كانت `config.fault` تساوي `true`.
    - تسجيل `Invalid phase detected` إذا كانت `duration <= 0`.
    - تسجيل `Switching to [color] for [duration] s` لكل مرحلة صالحة. استبدل `[color]` و`[duration]` بالخصائص المقابلة لكائن المرحلة.
    - على سبيل المثال، يجب أن تسجل `runSequence(config1, 1)` ما يلي:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. يجب أن يكون لديك دالة باسم `generateTimeline` بمعلمتين: `config` و `cycles`.

1. يجب أن تقوم الدالة `generateTimeline` بما يلي:
    - تسجيل الوقت التراكمي المنقضي بعد كل مرحلة عبر الدورات في مصفوفة، مع إضافة `duration` لكل مرحلة إلى المجموع الجاري أثناء التكرار.
    - معالجة جميع المراحل المعطوبة وغير الصالحة دون تحقق، حتى لو كانت قيمة `config.fault` تساوي `true` أو `duration <= 0`.
    - إرجاع مصفوفة الطوابع الزمنية التراكمية.
    - على سبيل المثال، يجب أن تُرجع `generateTimeline(config1, 1)` المصفوفة `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

لا يجب أن تجري تغييرات في كائنات `config1` و `config2` و `config3` و `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

يجب أن يكون لديك دالة باسم `runSequence`.

```js
assert.isFunction(runSequence);
```

يجب أن تقبل دالة `runSequence` معلمتين، `config` و `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

يجب أن تسجل `runSequence(config1, 1)` بالترتيب: `Switching to green for 5 s`، `Switching to yellow for 2 s`، `Switching to red for 4 s`.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

يجب أن تسجل `runSequence(config1, 2)` بالترتيب: `Switching to green for 5 s`، `Switching to yellow for 2 s`، `Switching to red for 4 s`، `Switching to green for 5 s`، `Switching to yellow for 2 s`، `Switching to red for 4 s`.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

يجب أن تسجل `runSequence(config2, 1)` بالترتيب: `Switching to red for 3 s`، `Invalid phase detected`، `Switching to green for 6 s`.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

يجب أن تسجل `runSequence(config3, 2)`: `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

يجب أن تسجل `runSequence(config4, 5)`: `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

يجب أن يكون لديك دالة باسم `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

يجب أن تقبل دالة `generateTimeline` معلمتين، `config` و `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

يجب أن تُرجع `generateTimeline(config1, 1)` المصفوفة `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

يجب أن تُرجع `generateTimeline(config1, 2)` المصفوفة `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

يجب أن تُرجع `generateTimeline(config2, 2)` المصفوفة `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

يجب أن تُرجع `generateTimeline(config3, 1)` المصفوفة `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

يجب أن تُرجع `generateTimeline(config4, 1)` المصفوفة `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
