---
id: 66c8ba41a77db20a93f9d7a1
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 JavaScript 조건문을 검토하며 대출 자격 확인기 앱을 만들어 봅니다.

이 앱은 사용자의 연간 소득과 신용 점수를 기준으로 복층 주택, 자동차, 콘도 대출 자격 여부를 확인합니다. 신용 점수는 빌린 돈을 얼마나 잘 관리하는지를 나타내는 숫자입니다.

시작하려면 다음 변수와 값을 만드세요.

| 변수 이름 | 값   |
| ----------- | ------- |
| `minIncomeForDuplex`       | `60000` |
| `minCreditScoreForDuplex`   | `700`   |
| `minIncomeForCondo`   | `45000`   |
| `minCreditScoreForCondo`    | `680`   |
| `minIncomeForCar`   | `30000`   |
| `minCreditScoreForCar`   | `650`   |

# --hints--

`minIncomeForDuplex` 변수를 `60000` 값으로 만드세요.

```js
assert.equal(minIncomeForDuplex, 60000);
```

`minCreditScoreForDuplex` 변수를 `700` 값으로 만드세요.

```js
assert.equal(minCreditScoreForDuplex, 700);
```

`minIncomeForCar` 변수를 `30000` 값으로 만드세요.

```js
assert.equal(minIncomeForCar, 30000);
```

`minCreditScoreForCar` 변수를 `650` 값으로 만드세요.

```js
assert.equal(minCreditScoreForCar, 650);
```

`minIncomeForCondo` 변수를 `45000` 값으로 만드세요.

```js
assert.equal(minIncomeForCondo, 45000);
```

`minCreditScoreForCondo` 변수를 `680` 값으로 만드세요.

```js
assert.equal(minCreditScoreForCondo, 680);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
