---
id: 66c8ba975ee7230e29f6c4ac
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

사용자가 대출 자격이 있을 때, 콘솔에 메시지를 표시하세요.

이를 위해, 신청자가 어떤 대출 자격이 있는지 반환하는 몇 가지 검사를 포함하는 함수를 작성할 것입니다.

빈 `getLoanMessage` 함수를 만들고 `annualIncome`과 `creditScore` 매개변수를 포함하세요.

# --hints--

`getLoanMessage` 함수를 만들어야 합니다.

```js
assert.isFunction(getLoanMessage)
```

`getLoanMessage` 함수는 `annualIncome`와 `creditScore`를 매개변수로 가져야 합니다.

```js
assert.match(getLoanMessage?.toString(), /\(\s*annualIncome,\s*creditScore\s*\)|\(\s*creditScore,\s*annualIncome\s*\)/);
```

`getLoanMessage`는 빈 함수여야 합니다.

```js
assert.match(getLoanMessage?.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

--fcc-editable-region--

--fcc-editable-region--
```
