---
id: 66c8ba975ee7230e29f6c4b2
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이제 `getLoanMessage` 함수를 테스트할 시간입니다.

아래 표를 사용해 `4` 변수들과 그 값을 만드세요:

| 변수 이름 | 값       |
| --------- | -------- |
| `duplexLoanMsg` | `getLoanMessage(85000, 850)` |
| `condoLoanMsg` | `getLoanMessage(65000, 690)` |
| `carLoanMsg` | `getLoanMessage(45000, 660)` |
| `noLoanMsg` | `getLoanMessage(25000, 550)` |

그다음, 각 변수를 콘솔에 출력해 메시지를 확인하세요.

이로써 대출 자격 확인기 프로젝트가 완성됩니다!

# --hints--

`duplexLoanMsg`이라는 변수가 있어야 합니다.

```js
assert.isNotNull(duplexLoanMsg);
```

`duplexLoanMsg` 변수는 `getLoanMessage(85000, 850)` 호출 결과로 설정해야 합니다.

```js
assert.equal(duplexLoanMsg, getLoanMessage(85000, 850));
```

`duplexLoanMsg` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(duplexLoanMsg\)/)
```

`condoLoanMsg`이라는 변수가 있어야 합니다.

```js
assert.isNotNull(condoLoanMsg);
```

`condoLoanMsg` 변수는 `getLoanMessage(65000, 690)` 호출 결과로 설정해야 합니다.

```js
assert.equal(condoLoanMsg, getLoanMessage(65000, 690));
```

`condoLoanMsg` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(condoLoanMsg\)/)
```

`carLoanMsg`이라는 변수가 있어야 합니다.

```js
assert.isNotNull(carLoanMsg);
```

`carLoanMsg` 변수는 `getLoanMessage(45000, 660)` 호출 결과로 설정해야 합니다.

```js
assert.equal(carLoanMsg, getLoanMessage(45000, 660));
```

`carLoanMsg` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(carLoanMsg\)/)
```

`noLoanMsg`이라는 변수가 있어야 합니다.

```js
assert.isNotNull(noLoanMsg);
```

`noLoanMsg` 변수는 `getLoanMessage(25000, 550)` 호출 결과로 설정해야 합니다.

```js
assert.equal(noLoanMsg, getLoanMessage(25000, 550));
```

`noLoanMsg` 변수를 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(noLoanMsg\)/)
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
  } else {
    return "You don't qualify for any loans."
  }
}

const duplexLoanMsg = getLoanMessage(85000, 850)
console.log(duplexLoanMsg)

const condoLoanMsg = getLoanMessage(65000, 690)
console.log(condoLoanMsg)

const carLoanMsg = getLoanMessage(45000, 660)
console.log(carLoanMsg)

const noLoanMsg = getLoanMessage(25000, 550)
console.log(noLoanMsg)
```
