---
id: 66c8ba975ee7230e29f6c4ad
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

`annualIncome`와 `creditScore`을 기준으로 사용자가 어떤 대출 자격이 있는지 확인하려면 `if/else if` 함수 안에 `getLoanMessage` 문장이나 삼항 연산자를 사용해야 합니다. 그런 다음 각 조건 블록에서 적절한 메시지를 반환하세요.

복층 대출부터 시작해 `annualIncome`가 `minIncomeForDuplex` 이상이고 `creditScore`가 `minCreditScoreForDuplex` 이상인지 확인하세요.

그 조건이 참이면 신청자는 복층 대출과 다른 대출 자격이 있습니다. 따라서 조건문 안에서 문자열 `"You qualify for a duplex, condo, and car loan."`을 반환하세요.

# --hints--

`getLoanMessage` 함수는 문자열을 반환해야 합니다.

```js
assert.isString(getLoanMessage(65000, 750));
assert.isString(getLoanMessage(60000, 750));
assert.isString(getLoanMessage(65000, 700));
assert.isString(getLoanMessage(60000, 700));
```

`getLoanMessage` 함수는 문자열 `"You qualify for a duplex, condo, and car loan."`을 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(65000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 750), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(65000, 700), "You qualify for a duplex, condo, and car loan.");
assert.strictEqual(getLoanMessage(60000, 700), "You qualify for a duplex, condo, and car loan.");
```

`getLoanMessage` 함수는 신청자의 연간 소득과 신용 점수가 복층 대출 요건을 충족하지 못하면 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(getLoanMessage(59000, 690));
assert.isUndefined(getLoanMessage(59000, 700));
assert.isUndefined(getLoanMessage(59000, 750));
assert.isUndefined(getLoanMessage(60000, 690));
assert.isUndefined(getLoanMessage(65000, 690));
```
 
# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
