---
id: 66c8ba975ee7230e29f6c4af
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

신청자의 연간 소득이 `minIncomeForCondo` 이상이고, 신용 점수가 `minCreditScoreForCondo` 이상이면 콘도와 자동차 대출 자격이 됩니다.

`getLoanMessage` 함수에서 이것이 참인지 확인하세요. 참이면 문자열 `"You qualify for a condo and car loan."`을 반환하세요.

# --hints--

`getLoanMessage` 함수는 신청자가 콘도와 자동차 대출 자격이 있으면 문자열을 반환해야 합니다.

```js
assert.isString(getLoanMessage(45000, 680));
```

`getLoanMessage` 함수는 문자열 `"You qualify for a condo and car loan."`을 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(45000, 680), "You qualify for a condo and car loan.");
```

`getLoanMessage` 함수는 신청자의 연간 소득과 신용 점수가 콘도 대출 요건을 충족하지 못하면 `undefined`을 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(45000, 650), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
