---
id: 66c8ba975ee7230e29f6c4b1
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

신청자의 연간 소득과 신용 점수가 `minIncomeForCar`와 `minCreditScoreForCar`보다 낮으면 어떤 대출도 받을 자격이 없습니다. 따라서 문자열 `"You don't qualify for any loans."`를 반환하세요.

# --hints--

신청자가 어떤 대출 자격도 갖추지 못하면 `getLoanMessage` 함수가 문자열을 반환해야 합니다.

```js
assert.isString(getLoanMessage(25000, 550));
```

신청자가 어떤 대출 요건도 충족하지 못하면 `getLoanMessage` 함수가 문자열 `"You don't qualify for any loans."`를 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(15000, 600), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(25000, 550), "You don't qualify for any loans.");
assert.strictEqual(getLoanMessage(20000, 500), "You don't qualify for any loans.");
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
  } else if (annualIncome >= minIncomeForCar && creditScore >= minCreditScoreForCar) {
    return "You qualify for a car loan."
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
