---
id: 66c8ba975ee7230e29f6c4b0
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 신청자가 자동차 대출 자격이 있는지 확인하세요. 자격이 있다면 문자열 `"You qualify for a car loan."`을 반환하세요.

# --hints--

`getLoanMessage` 함수는 신청자가 자동차 대출 자격이 있을 때 문자열을 반환해야 합니다.

```js
assert.isString(getLoanMessage(30000, 650));
```

`getLoanMessage` 함수는 문자열 `"You qualify for a car loan."`을 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(30000, 650), "You qualify for a car loan.");
```

`getLoanMessage` 함수는 신청자의 연간 소득과 신용 점수가 자동차 대출 요건을 충족하지 못하면 `undefined`를 반환해야 합니다.

```js
assert.strictEqual(getLoanMessage(30000, 600), undefined);
```

# --seed--

## --seed-contents--

```js
const minIncomeForDuplex = 60000;
const minCreditScoreForDuplex = 700;

const minIncomeForCondo = 45000;
const minCreditScoreForCondo = 680;

const minIncomeForCar = 30000;
const minCreditScoreForCar = 650;

function getLoanMessage(annualIncome, creditScore) {
  if(creditScore >= minCreditScoreForDuplex && annualIncome >= minIncomeForDuplex) {
    return "You qualify for a duplex, condo, and car loan."
  } else if (annualIncome >= minIncomeForCondo && creditScore >= minCreditScoreForCondo) {
    return "You qualify for a condo and car loan."
--fcc-editable-region--
  } 
  
--fcc-editable-region--
}
```
