---
id: ab6137d4e35944e21037b769
title: Створіть конвертер рядка у формат title case
challengeType: 26
dashedName: build-a-title-case-converter
---

# --description--

У цій лабораторній роботі ви створите функцію, яка перетворює рядок у формат title case. Title case означає, що перша літера кожного слова є великою, а решта слова — у нижньому регістрі.

`"Web Development Is Awesome"` — це приклад рядка у форматі title case.

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні мати `titleCase` функцію, яка приймає рядок як аргумент.
1. `titleCase` функція повинна повертати рядок, у якому перша літера кожного слова велика, а решта слова — у нижньому регістрі.
1. `titleCase("I like to code")` має повертати `"I Like To Code"`.
1. `titleCase("javaScript is fun")` має повертати `"Javascript Is Fun"`.

# --hints--

Ви повинні створити функцію з назвою `titleCase`.

```js
assert.isFunction(titleCase);
```

`titleCase` має приймати один аргумент.

```js
assert.lengthOf(titleCase, 1);
```


`titleCase("I'm a little tea pot")` має повертати рядок.

```js
assert.isString(titleCase("I'm a little tea pot"));
```

`titleCase("I'm a little tea pot")` має повертати рядок `I'm A Little Tea Pot`.

```js
assert.strictEqual(titleCase("I'm a little tea pot"), "I'm A Little Tea Pot");
```

`titleCase("sHoRt AnD sToUt")` має повертати рядок `Short And Stout`.

```js
assert.strictEqual(titleCase('sHoRt AnD sToUt'), 'Short And Stout');
```

`titleCase("HERE IS MY HANDLE HERE IS MY SPOUT")` має повертати рядок `Here Is My Handle Here Is My Spout`.

```js
assert.strictEqual(
  titleCase('HERE IS MY HANDLE HERE IS MY SPOUT'),
  'Here Is My Handle Here Is My Spout'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function titleCase(str) {
  return str
    .split(' ')
    .map(word => word.charAt(0).toUpperCase() + word.substring(1).toLowerCase())
    .join(' ');
}
```
