---
id: 69f0a37c3bf194a7c9c62492
title: Crear un libro mayor de préstamos de dispositivos
challengeType: 26
dashedName: lab-device-loan-ledger
---

# --description--

Eres un desarrollador de software encargado de crear una aplicación de libro mayor para gestionar los dispositivos de hardware de TI proporcionados por el Service Desk de tu empresa. El libro mayor es un objeto cuyas claves son las etiquetas de los recursos y cuyos valores son objetos:

```js
const ledger = {
  "1": {
    type: "Laptop",
    status: "CheckedIn",
    borrower: {
      name: "",
      email: ""
    },
    dueDate: ""
  }
};
```

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes modelar el libro mayor como un objeto donde cada dispositivo esté identificado por su etiqueta de recurso y tenga las propiedades `type`, `status`, `borrower` (con `name` y `email`) y `dueDate`.

1. Debes implementar una función `checkoutDevice(ledger, assetTag, borrower)` que reciba un objeto libro mayor y lo clone sin modificar el original. En el clon, debe establecer el `name` y `email` del objeto `borrower`, actualizar el `status` a `"CheckedOut"` y devolver un objeto con la forma `{ ledger: updatedLedger, message: confirmationString }`, donde el `message` incluye la etiqueta del recurso y el `name` del objeto `borrower`.

   - Si `assetTag` no existe en el libro mayor, la función debe devolver el libro mayor sin cambios junto con un `message` que incluya la etiqueta del recurso y explique que no se encontró.
   - Si el dispositivo ya está `"CheckedOut"`, la función debe devolver el libro mayor sin cambios y mantener al prestatario actual, junto con un `message` que incluya la etiqueta del recurso y explique que el dispositivo ya está prestado.

1. Debes implementar una función `checkinDevice(ledger, assetTag)` que reciba un objeto libro mayor y lo clone sin modificar el original. En el clon, debe borrar el `name` y `email` del objeto `borrower`, reiniciar `dueDate` a `""`, actualizar el `status` a `"CheckedIn"` y devolver un objeto con la forma `{ ledger: updatedLedger, message: confirmationString }`, donde el `message` incluye la etiqueta del recurso.

   - Si `assetTag` no existe en el libro mayor, la función debe devolver el libro mayor sin cambios junto con un `message` que incluya la etiqueta del recurso y explique que no se encontró.

1. Debes implementar una función `listOverdueDevices(ledger, today)` que reciba un objeto libro mayor y devuelva un arreglo con los objetos de dispositivos que están atrasados, ordenados por `dueDate` en orden ascendente. Un dispositivo está atrasado cuando su `status` es `"CheckedOut"` y su `dueDate` es estrictamente anterior a `today`. Tanto `today` como `dueDate` son cadenas en orden mes/día/año, y el mes y día pueden o no tener ceros a la izquierda (por ejemplo, `"9/5/2025"` y `"09/05/2025"` son válidos).

1. No debes usar el objeto `Date` de JavaScript en ninguna parte de tu código.

1. Debes implementar una función `serializeLedger(ledger)` que convierta el objeto libro mayor en una cadena JSON, y una función `loadLedger(json)` que convierta una cadena JSON de vuelta a un objeto JavaScript.

# --hints--

Debes crear una función llamada `checkoutDevice`.

```js
assert.isFunction(checkoutDevice);
```

Debes crear una función llamada `checkinDevice`.

```js
assert.isFunction(checkinDevice);
```

Debes crear una función llamada `listOverdueDevices`.

```js
assert.isFunction(listOverdueDevices);
```

Debes crear una función llamada `serializeLedger`.

```js
assert.isFunction(serializeLedger);
```

Debes crear una función llamada `loadLedger`.

```js
assert.isFunction(loadLedger);
```

Tu función `checkoutDevice` no debe modificar el libro mayor original.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(ledger, original);
```

Tu función `checkoutDevice` debe establecer el `name` y `email` del prestatario y poner el `status` en `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.isObject(result);
assert.equal(result.ledger["1"].status, "CheckedOut");
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
```

Tu función `checkoutDevice` debe devolver un mensaje de confirmación que incluya la etiqueta del recurso y el `name` del prestatario.

```js
const ledger = {
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "2", { name: "Bob", email: "bob@test.com" });
assert.isString(result.message);
assert.include(result.message, "2");
assert.include(result.message, "Bob");
```

Tu función `checkoutDevice` debe preservar los dispositivos no relacionados y los campos sin cambios en el libro mayor devuelto.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "1", { name: "Alice", email: "alice@test.com" });
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Tu función `checkoutDevice` debe devolver el libro mayor sin cambios y un mensaje que incluya la etiqueta del recurso cuando la etiqueta no se encuentre.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkoutDevice(ledger, "99", { name: "Alice", email: "alice@test.com" });
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Tu función `checkoutDevice` debe devolver el libro mayor sin cambios, mantener al prestatario actual y devolver un mensaje que incluya la etiqueta del recurso cuando el dispositivo ya esté prestado.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkoutDevice(ledger, "1", { name: "Bob", email: "bob@test.com" });
assert.equal(result.ledger["1"].borrower.name, "Alice");
assert.equal(result.ledger["1"].borrower.email, "alice@test.com");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "1");
```

Tu función `checkinDevice` no debe modificar el libro mayor original.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const original = JSON.parse(JSON.stringify(ledger));
checkinDevice(ledger, "1");
assert.deepEqual(ledger, original);
```

Tu función `checkinDevice` debe borrar el `name` y `email` del prestatario, reiniciar `dueDate` a `""` y poner el `status` en `"CheckedIn"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].status, "CheckedIn");
assert.equal(result.ledger["1"].borrower.name, "");
assert.equal(result.ledger["1"].borrower.email, "");
assert.equal(result.ledger["1"].dueDate, "");
```

Tu función `checkinDevice` debe devolver un mensaje de confirmación que incluya la etiqueta del recurso.

```js
const ledger = {
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Carol", email: "carol@test.com" }, dueDate: "11/15/2025" }
};
const result = checkinDevice(ledger, "3");
assert.isString(result.message);
assert.include(result.message, "3");
```

Tu función `checkinDevice` debe preservar los dispositivos no relacionados y los campos sin cambios en el libro mayor devuelto.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
const result = checkinDevice(ledger, "1");
assert.equal(result.ledger["1"].type, "Laptop");
assert.deepEqual(result.ledger["2"], ledger["2"]);
```

Tu función `checkinDevice` debe devolver el libro mayor sin cambios y un mensaje que incluya la etiqueta del recurso cuando la etiqueta no se encuentre.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "12/01/2025" }
};
const result = checkinDevice(ledger, "99");
assert.deepEqual(result.ledger, ledger);
assert.isString(result.message);
assert.include(result.message, "99");
```

Tu función `listOverdueDevices` debe devolver un arreglo.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};
assert.isArray(listOverdueDevices(ledger, "12/01/2025"));
```

Tu función `listOverdueDevices` debe devolver solo los dispositivos atrasados con un `status` de `"CheckedOut"`.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "10/01/2025" },
  "2": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "12/31/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 1);
assert.equal(result[0].borrower.name, "Alice");
```

Tu función `listOverdueDevices` debe devolver los dispositivos atrasados ordenados por `dueDate` en orden ascendente.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "11/01/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/01/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Bob");
assert.equal(result[1].borrower.name, "Alice");
```

Tu función `listOverdueDevices` debe ordenar correctamente incluso cuando el mes o el día no tienen ceros a la izquierda.

```js
const ledger = {
  "1": { type: "Laptop", status: "CheckedOut", borrower: { name: "Alice", email: "alice@test.com" }, dueDate: "9/5/2025" },
  "2": { type: "PC", status: "CheckedOut", borrower: { name: "Bob", email: "bob@test.com" }, dueDate: "10/1/2025" }
};
const result = listOverdueDevices(ledger, "12/01/2025");
assert.lengthOf(result, 2);
assert.equal(result[0].borrower.name, "Alice");
assert.equal(result[1].borrower.name, "Bob");
```

No debes usar el objeto `Date` de JavaScript en ninguna parte de tu código.

```js
assert.notMatch(__helpers.removeJSComments(code), /\bnew\s+Date\b|\bDate\s*[.(]/);
```

Tu función `serializeLedger` debe devolver una cadena JSON que represente el libro mayor que se le pasó.

```js
const obj = { "1": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" } };
assert.isString(serializeLedger(obj));
assert.deepEqual(JSON.parse(serializeLedger(obj)), obj);
```

Tu función `loadLedger` debe devolver el objeto JavaScript que representa la cadena JSON.

```js
const json = '{"1":{"type":"Laptop","status":"CheckedIn","borrower":{"name":"","email":""},"dueDate":""}}';
assert.isObject(loadLedger(json));
assert.deepEqual(loadLedger(json), JSON.parse(json));
```

# --seed--

## --seed-contents--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

```

# --solutions--

```js
const equipmentLedger = {
  "1": { type: "PC", status: "CheckedOut", borrower: { name: "John Smith", email: "john@acme.org" }, dueDate: "11/30/2025" },
  "2": { type: "Laptop", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" },
  "3": { type: "Laptop", status: "CheckedOut", borrower: { name: "Jane Doe", email: "jane@acme.org" }, dueDate: "10/31/2025" },
  "4": { type: "iPad", status: "CheckedIn", borrower: { name: "", email: "" }, dueDate: "" }
};

function serializeLedger(ledger) {
  return JSON.stringify(ledger);
}

function loadLedger(json) {
  return JSON.parse(json);
}

function checkoutDevice(ledger, assetTag, borrower) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  if (updatedLedger[assetTag].status === "CheckedOut") {
    return {
      ledger: updatedLedger,
      message: `${updatedLedger[assetTag].type} with asset tag ${assetTag} is already checked out.`
    };
  }

  updatedLedger[assetTag].status = "CheckedOut";
  updatedLedger[assetTag].borrower.name = borrower.name;
  updatedLedger[assetTag].borrower.email = borrower.email;
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked out to ${borrower.name}.`;
  return { ledger: updatedLedger, message };
}

function checkinDevice(ledger, assetTag) {
  const updatedLedger = JSON.parse(JSON.stringify(ledger));

  if (!updatedLedger[assetTag]) {
    return {
      ledger: updatedLedger,
      message: `Asset tag ${assetTag} was not found in the ledger.`
    };
  }

  updatedLedger[assetTag].status = "CheckedIn";
  updatedLedger[assetTag].borrower.name = "";
  updatedLedger[assetTag].borrower.email = "";
  updatedLedger[assetTag].dueDate = "";
  const message = `${updatedLedger[assetTag].type} with asset tag ${assetTag} has been checked in.`;
  return { ledger: updatedLedger, message };
}

function reformatDate(date) {
  const datePartsArray = date.split("/");
  let month = datePartsArray[0];
  let day = datePartsArray[1];
  const year = datePartsArray[2];

  month = month.padStart(2, "0");
  day = day.padStart(2, "0");

  const formattedDate = year + month + day;
  return parseInt(formattedDate, 10);
}

function listOverdueDevices(ledger, today) {
  const todayFormatted = reformatDate(today);
  const overdue = [];

  for (const assetTag in ledger) {
    const item = ledger[assetTag];
    if (item.status === "CheckedOut" && item.dueDate) {
      if (reformatDate(item.dueDate) < todayFormatted) {
        overdue.push(item);
      }
    }
  }

  return overdue.sort((a, b) => reformatDate(a.dueDate) - reformatDate(b.dueDate));
}
```
