---
id: 694bf6eb530e19a9c48c59f4
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`+`를 사용해 여러 문자열을 연결하고 `str()`로 숫자를 변환하는 것은 복잡하고 읽기 어려울 수 있습니다.

Python 3.6에서는 이를 해결하기 위해 **f-strings**를 도입했습니다. 따옴표 앞에 `f` 문자를 추가하면, 중괄호 `{}`로 나타낸 치환 필드 안에 변수와 표현식을 넣을 수 있습니다. 예를 들어:

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

변수 `employee_card`를 만들고, `Employee:` 다음에 공백과 `full_name` 변수 값을 표시하는 f-string을 할당하세요.

# --hints--

`employee_card`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

`employee_card` 변수는 `f'Employee: {full_name}'` 값을 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
