---
id: 694bf3af05d10fd89dab05cc
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

Python에서는 `+` 연산자를 사용해 문자열을 결합할 수 있습니다. 이를 연결(concatenation)이라고 합니다. 예시는 다음과 같습니다:

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

`full_name`와 `first_name`를 연결하여 `last_name` 변수를 만드세요. 그런 다음 `full_name`를 출력하세요.

# --hints--

`full_name`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

`first_name`와 `last_name`를 연결해 만든 문자열을 `full_name` 변수에 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

`full_name` 변수를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
