---
id: 694bf6eb530e19a9c48c59f6
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

문자열의 시작 부분뿐만 아니라 어느 부분에서든 슬라이스할 수 있습니다. 그리고 이것은 여러 경우에 유용합니다.

변수 `year_code`를 만들고 `employee_code`의 인덱스 `4`부터 `8`까지의 슬라이스를 할당하세요. 이렇게 하면 `2026`를 추출합니다.

그다음 변수 `initials`를 만들고 `employee_code`의 인덱스 `9`부터 `11`까지의 슬라이스를 할당하세요. 이렇게 하면 `JD`를 추출합니다.

마지막으로 두 변수를 터미널에 출력하세요.

# --hints--

`year_code`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

`employee_code`의 인덱스 `4`부터 `8`까지 슬라이스해서 결과를 `year_code` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

`initials`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

`employee_code`의 인덱스 `9`부터 `11`까지 슬라이스해서 결과를 `initials` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

`year_code`를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

`initials`를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
