---
id: 694bf6eb530e19a9c48c59f5
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

문자열을 다룰 때 특정 부분을 추출해야 하는 경우가 많습니다. 이를 **슬라이싱**이라고 합니다.

구문은 `string[start:stop]`이며, 여기서:

* `start`은 슬라이스가 시작되는 인덱스(**포함**)입니다.
* `stop`는 슬라이스가 끝나는 인덱스(**미포함**)입니다.

예를 들어, `text = 'Python'`라면 `text[0:2]`는 `'Py'`를 반환합니다.

`employee_code`를 `'DEV-2026-JD-001'`로 정의하세요. 그다음 `department` 변수를 만들고 `employee_code`의 인덱스 `0`부터 `3`까지 슬라이스한 값을 할당하세요. 그리고 `department`를 터미널에 출력하세요.

# --hints--
`employee_code`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

`DEV-2026-JD-001` 문자열을 `employee_code` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

`department`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

`employee_code`의 처음 세 글자를 슬라이스하여 `department` 변수에 할당해야 합니다.

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

`department`를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
