---
id: 694bf6ea530e19a9c48c59f0
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

`'John' + 'Doe'`처럼 두 문자열을 이어 붙이면 공백 없이 `'JohnDoe'`가 됩니다. 이를 고치려면 두 문자열 사이에 공백이 포함된 문자열(`' '`)을 이어 붙여야 합니다.

`full_name` 변수를 업데이트해서 `first_name`, 공백, 그리고 `last_name`를 이어 붙이도록 하세요.

# --hints--

`' '`와 `first_name` 사이에 공백(`last_name`)을 추가해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
