---
id: 694bf6eb530e19a9c48c59f1
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이제 주소가 불완전해 보입니다. 직원이 사는 아파트 번호도 추가하고 싶으니 변수를 수정해야 합니다.  

기존 문자열 변수 끝에 내용을 추가하려면 **복합 할당** 연산자 `+=`를 사용할 수 있습니다. 이는 `var = var + 'new text'`를 쓰는 것보다 더 간단합니다. 예를 들어:  

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

문자열은 불변이므로 이 연산은 원래 문자열을 변경하지 않습니다. 대신 새 문자열을 만들고 변수를 다시 할당합니다.

`+=` 연산자를 사용해 `, Apartment 4B` 변수에 `address` 문자열을 추가하세요.

# --hints--

`+=` 변수에 `, Apartment 4B` 문자열을 추가하려면 `address` 연산자를 사용하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
