---
id: 694bf6eb530e19a9c48c59f3
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 `str()` 함수를 한 번 더 사용할 차례입니다. 나이와 마찬가지로, 숫자형 변수를 다른 텍스트와 연결하기 전에 반드시 문자열로 변환해야 합니다.

`experience_years`이라는 변수를 만들고 정수 `5`를 할당하세요.

그다음, `experience_info` 변수를 만드세요. `'Experience: '`, `experience_years` 변수(문자열로 변환됨), `' years'`을 연결해 만든 문자열을 할당하세요. 결과를 터미널에 출력하세요.

# --hints--

`experience_years`이라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

정수 `5`를 `experience_years` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

`experience_info`이라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

`'Experience: '`, `str(experience_years)`, `' years'`을 연결해 만든 문자열을 `experience_info` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

`experience_info`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
