---
id: 694bf6eb530e19a9c48c59f7
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

문자열 끝에서부터 슬라이스할 때 음수를 사용할 수도 있습니다. 예를 들어, `-1`는 마지막 문자를 나타내고, `-2`는 끝에서 두 번째 문자를 나타내는 식입니다.

문자열의 마지막 세 문자를 얻으려면 `string[-3:]`를 사용할 수 있습니다. 콜론 뒤에 stop 매개변수가 생략된 점에 주목하세요. 이는 슬라이스가 문자열 경계까지 진행된다는 뜻입니다.

`last_three`라는 변수를 만드세요. 음수 인덱싱을 사용해 `employee_code`에서 마지막 세 문자를 추출하세요(이는 ID 번호를 나타냅니다). 마지막으로 `last_three`를 터미널에 출력하세요.

이로써 직원 카드 워크숍이 완료됩니다.

# --hints--

`last_three`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

음수 인덱싱을 사용해 `employee_code`의 마지막 세 문자를 슬라이스하고 `last_three` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

`last_three`를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
