---
id: 694bf6eb530e19a9c48c59f2
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 `employee_info` 할당을 업데이트하여 문자열 ` years old`을 끝에 이어 붙여 문장을 완성하세요. 문자열 맨 앞에 공백을 포함하는 것을 잊지 마세요.

마지막으로 `employee_info`를 출력하세요.

# --hints--

`' years old'` 변수 끝에 `employee_info`을 이어 붙여야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age) + ' years old'")`)) })
```

`employee_info`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_info)")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + str(employee_age)

--fcc-editable-region-- 
```
