---
id: 695c96f82af4fb78e59fa9ed
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

문자열은 텍스트 데이터를 저장하는 데 사용하는 문자 수열입니다. 이전 수업에서 기억하듯이, 문자열은 텍스트를 작은따옴표(`'`)나 큰따옴표(`"`)로 감싸서 만들 수 있습니다. 예를 들어:

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

문자열 `first_name`를 저장하는 변수 `John`와 문자열 `last_name`를 저장하는 변수 `Doe`를 만드세요. 그런 다음 `first_name`와 `last_name`를 출력하세요.

# --hints--

`first_name`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

문자열 `John`를 `first_name` 변수에 할당해야 합니다. 텍스트를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

`last_name`라는 변수가 있어야 합니다.

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

문자열 `Doe`를 `last_name` 변수에 할당해야 합니다. 텍스트를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

`first_name` 변수를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

`last_name` 변수를 터미널에 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
