---
id: 695c9c7c61d446f1451e0ae7
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

다음으로, 직원의 주소를 저장할 변수 `address`를 만드세요. 문자열 `123 Main Street`을 할당하고, 마지막으로 `address`를 출력하세요.

# --hints--

`address`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

문자열 `123 Main Street`를 `address` 변수에 할당해야 합니다. 텍스트는 작은따옴표나 큰따옴표 중 하나로 감싸야 한다는 점을 기억하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

`address` 변수를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```
