---
id: 695c9f21a8b1d53cf1b97841
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 `employee_age`라는 변수를 만들고 정수 `28`을 할당하세요.


# --hints--

`employee_age`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

정수 `28`을 `employee_age` 변수에 할당하세요. 텍스트를 따옴표로 감싸지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
