---
id: 695df6480de23758af5a29cd
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

이제 직원의 나이를 표시하는 문자열을 만들고 싶습니다.

변수 `employee_info`를 만들고 다음을 연결한 결과를 할당하는 것부터 시작하세요:

- `full_name` 변수
- 앞뒤로 공백이 있는 `is` 문자로 이루어진 문자열

# --hints--

`employee_info`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

`full_name + ' is '`를 `employee_info` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
