---
id: 695ea1e55c0e619e597fc5c9
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

보시다시피 Python이 `TypeError: can only concatenate str (not "int") to str`를 발생시켰습니다. 이는 Python이 텍스트(문자열)와 숫자(정수)를 직접 연결하는 것을 허용하지 않기 때문입니다.

이를 고치려면 먼저 `str()` 함수를 사용해 숫자를 문자열로 변환해야 합니다. 이 함수는 객체의 문자열 버전을 반환합니다.

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

`employee_info`를 사용해 `employee_age`를 문자열로 변환하도록 `str(employee_age)` 할당을 업데이트하세요.

# --hints--

`employee_info`를 사용해 `employee_age`를 문자열로 변환하여 `str(employee_age)` 변수의 할당을 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
