---
id: 695fd3dde03273875e5d6709
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`employee_info` 할당을 업데이트하여 `employee_age`도 끝에 이어 붙이세요.

이 작업을 완료하면 터미널에 `TypeError`가 표시됩니다. 다음 단계에서 이를 고치는 작업을 하게 됩니다.

# --hints--

`employee_info` 할당을 업데이트하여 `employee_age`를 끝에 이어 붙이세요.

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
