---
id: 695fcfbba648b3a96a9495a4
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

터미널 출력이 점점 복잡해지고 있습니다. 계속하기 전에 정리할 시간입니다.

코드에서 모든 `print()` 문장을 제거하세요.

# --hints--

코드에 `print` 호출이 없어야 합니다.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
