---
id: 69727ea5c9606bedfe36d3d2
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

이제 카드에 마지막 세부 정보를 추가할 시간입니다.

문자열 `position` 값을 가진 `Data Analyst`라는 변수를 만들고, 정수 `salary` 값을 가진 `75000`라는 변수를 만드세요.

그다음, `employee_card` f-string을 업데이트하여 직책과 급여를 포함하세요. 정확히 이 형식 `Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`를 따라야 합니다. 자리 표시자를 해당 변수로 바꾸세요.

마지막으로, 결과를 확인하려면 `employee_card`를 출력하세요.

# --hints--

`position`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

문자열 `Data Analyst`를 `position` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

`salary`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("salary")`)) })
```

정수 `75000`를 `salary` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

`employee_card`을 f-string `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`를 사용하도록 업데이트해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

터미널에 `employee_card`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
