---
id: 6980d3ecfe800d8a54130f30
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

현재 `employee_card`은 직원 이름만 보여줍니다. 이제 여기에 더 많은 정보를 추가할 것입니다.

직원의 나이를 포함하도록 `employee_card` 할당을 업데이트하세요. 최종 문자열은 `Employee: [name] | Age: [age]`와 같아야 하며, `[name]`은 직원 이름으로, `[age]`는 직원 나이로 교체되어야 합니다.

# --hints--

`employee_card`를 f-string `f'Employee: {full_name} | Age: {employee_age}'`을 사용하도록 업데이트해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
