---
id: 60f02e7361b68405e27b62a5
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本工作坊，你将通过设计一个注册表单来学习如何设置表单样式。

所有 HTML 模板（`DOCTYPE`、`html`、`head` 和 `body`）已为你提供。

在 `body` 中，使用文本为 `Registration Form` 的 `h1` 元素添加一个标题，为表单提供上下文信息。

# --hints--

你应该在 `body` 中添加一个 `h1`。

```js
assert.exists(document.querySelector('body > h1'));
```

你应该给 `h1` 元素添加文本 `Registration Form`。

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
