---
id: 60f1922fcbd2410527b3bd89
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

In previous lessons, you learned how to work with viewport units.

记住，`vh` 单位代表视口高度，等于视口 `height` 的 1%。 这使它与视口高度相对应。

现在是时候用 CSS 来装饰项目了。 首先设置 `body` 的 `width` 为 `100%`，`height` 为 `100vh`。

# --hints--

你应该使用 `body` 元素选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

你应该给 `body` 一个 `width`，设置为 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

你应该给 `body` 一个 `height`，设置为 `100vh`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```

