---
id: 60f81616cff80508badf9ad5
title: 步骤 16
challengeType: 0
dashedName: step-16
---

# --description--

此时，就可以提交表单了。 不过，你可能会注意到没什么变化。

为使表单更具交互性，在第一个 `fieldset` 的 `input` 元素中添加 `required` 属性。

现在，如果你在未填写必填字段的情况下尝试提交表单，就会看到一条错误信息。

# --hints--

你应该给第一个 `input` 元素添加一个 `required` 属性。

```js
assert.isTrue(document.querySelector('input')?.required);
```

你应该给第二个 `input` 元素添加一个 `required` 属性。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.required);
```

你应该给第三个 `input` 元素添加一个 `required` 属性。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.required);
```

你应该给第四个 `input` 元素添加一个 `required` 属性。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.required);
```

你不应给 `submit` 类型的 `input` 添加 `required` 属性。

```js
assert.isFalse(document.querySelector('input[type="submit"]')?.required);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
