---
id: 60f81167d0d4910809f88945
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

`type` 为 `submit` 的第一个 `input` 元素会自动设置为提交其最近的 `form` 父元素。

为了处理表格提交，在最后一个 `fieldset` 元素之后添加一个 `input` 元素，`type` 属性设置为 `submit`，`value` 属性设置为 `Submit`。

# --hints--

你应该在最后一个 `fieldset` 元素后添加 `input` 元素。

```js
assert.exists(document.querySelectorAll('fieldset')?.[2]?.nextElementSibling?.tagName, 'input');
```

你应该给 `input` 元素一个 `type` 属性，值为 `submit`。

```js
assert.exists(document.querySelector('fieldset + input[type="submit"]'));
```

你应该给 `input` 元素一个 `value` 属性，值为 `Submit`。

```js
assert.exists(document.querySelector('fieldset + input[value="Submit"]'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
    </fieldset>
    <fieldset></fieldset>
    <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
