---
id: 60f8618d191b940d62038513
title: 步骤 22
challengeType: 0
dashedName: step-22
---

# --description--

你想每次只能选择一个单选输入。 但是，表单并不知道这些单选输入是相关的。

给每一个单选按钮添加相同的 `name` 属性并设置其值为 `account-type`，以关联单选按钮。 现在，无法同时选择两个单选输入了。

# --hints--

你应该为第一个单选输入添加 `name` 属性，值为 `account-type`。

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(1) input[type="radio"]')?.name, 'account-type');
```

你应该为第二个单选输入添加 `name` 属性，值为 `account-type`。

```js
assert.equal(document.querySelector('fieldset:nth-child(2) label:nth-child(2) input[type="radio"]')?.name, 'account-type');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
--fcc-editable-region--
      <label><input type="radio" /> Personal</label>
      <label><input type="radio" /> Business</label>
--fcc-editable-region--
    </fieldset>
    <fieldset></fieldset>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
