---
id: 60fabf0dd4959805dbae09e6
title: 步骤 29
challengeType: 0
dashedName: step-29
---

# --description--

在第一个 `label` 之后添加另一个，文本为 `Input your age (years): `。然后，嵌套一个 `type` 为 `number` 的 `input`。

接下来，在 `input` 中添加一个 `min` 属性，其值为 `13`，因为 13 岁以下的用户不应注册。 此外，用户的年龄可能不会超过 120 岁；添加一个值为 `120` 的 `max` 属性。

现在，如果有人试图用超出范围的值提交表单，就会出现警告，表单将无法提交。 试一下吧。

# --hints--

你应在第三个 `fieldset` 中，在现有 `label` 之后再添加一个 `label`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label + label'));
```

你应该给 `label` 赋予文本 `Input your age (years): `。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2)')?.textContent?.trim(), 'Input your age (years):');
```

你应该给 `label` 一个 `input`，`type` 为 `number`。

```js
assert.exists(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]'));
```

你应该给 `input` 一个 `min` 属性，值为 `13`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.min, '13');
```

你应该给 `input` 一个 `max` 属性，值为 `120`。

```js
assert.equal(document.querySelector('fieldset:nth-child(3) > label:nth-child(2) > input[type="number"]')?.max, '120');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label>Upload a profile picture: <input type="file" /></label>
--fcc-editable-region--
      
--fcc-editable-region--
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
