---
id: 60fad1cafcde010995e15306
title: 步骤 38
challengeType: 0
dashedName: step-38
---

# --description--

在提交表单时，为每个可提交元素提供 `name` 属性是非常有用的，也是好的实践。 该属性用于标识表单提交中的元素。

除了两个 `radio` 输入（你已经命名的）外，为每个提交的元素指定唯一的 `name` 属性。

# --hints--

你应该为期望第一个名字的 `input` 提供一个 `name` 属性。*附注 我本会选择 `first-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(1) > input')?.name);
```

你应该为期望姓氏的 `input` 提供一个 `name` 属性。*附注 我本来会选择 `last-name`*

```js
assert.isNotEmpty(document.querySelector('fieldset:nth-child(1) > label:nth-child(2) > input')?.name);
```

你应该为 `email` 添加一个 `name` 属性。*附注 我本来会选择 `email`*

```js
assert.isNotEmpty(document.querySelector('input[type="email"]')?.name);
```

你应该给 `password` 一个 `name` 属性。*附注 我本来会选择 `password`*

```js
assert.isNotEmpty(document.querySelector('input[type="password"]')?.name);
```

你应该为 `checkbox` 添加一个 `name` 属性。*附注 我本来会选择 `terms`*

```js
assert.isNotEmpty(document.querySelector('input[type="checkbox"]')?.name);
```

你应该为 `file` 赋予一个 `name` 属性。*附注 我本来会选择 `file`*

```js
assert.isNotEmpty(document.querySelector('input[type="file"]')?.name);
```

你应该给 `number` 一个 `name` 属性。*附注 我本来会选择 `age`*

```js
assert.isNotEmpty(document.querySelector('input[type="number"]')?.name);
```

你应该为 `select` 元素提供一个 `name` 属性。*附注 我本来会选择 `referrer`*

```js
assert.isNotEmpty(document.querySelector('select')?.name);
```

你应该为 `textarea` 元素添加一个 `name` 属性。*附注 我本来会选择 `bio`*

```js
assert.isNotEmpty(document.querySelector('textarea')?.name);
```

你不该给任何 `option` 元素添加 `name` 属性。

```js
[...document.querySelectorAll('option')]?.forEach(option => assert.isUndefined(option?.name));
```

你不该给任何 `label` 元素添加 `name` 属性。

```js
[...document.querySelectorAll('label')]?.forEach(label => assert.isUndefined(label?.name));
```

你不该给任何 `fieldset` 元素添加 `name` 属性。

```js
[...document.querySelectorAll('fieldset')]?.forEach(fieldset => assert.isEmpty(fieldset?.name));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Registration Form</title>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <h1>Registration Form</h1>
  <p>Please fill out this form with the required information</p>
  <form method="post" action='https://register-demo.freecodecamp.org'>
    <fieldset>
--fcc-editable-region--
      <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
      <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
      <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
      <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
    </fieldset>
    <fieldset>
      <legend>Account type (required)</legend>
      <label for="personal-account"><input id="personal-account" type="radio" name="account-type" value="personal" checked /> Personal</label>
      <label for="business-account"><input id="business-account" type="radio" name="account-type" value="business" /> Business</label>
    </fieldset>
    <fieldset>
      <label for="profile-picture">Upload a profile picture: <input id="profile-picture" type="file"/></label>
      <label for="age">Input your age (years): <input id="age" type="number" min="13" max="120" /></label>
      <label for="referrer">How did you hear about us?
        <select id="referrer">
          <option value="">(select one)</option>
          <option value="1">freeCodeCamp News</option>
          <option value="2">freeCodeCamp YouTube Channel</option>
          <option value="3">freeCodeCamp Forum</option>
          <option value="4">Other</option>
        </select>
      </label>
      <label for="bio">Provide a bio:
        <textarea id ="bio" rows="3" cols="30" placeholder="I like coding on the beach..."></textarea>
      </label>
    </fieldset>
    <a href="https://www.freecodecamp.org/news/terms-of-service">Read our terms and conditions</a>
    <label for="terms-and-conditions">
      <input id="terms-and-conditions" type="checkbox" required />
      I accept the terms and conditions
    </label>
--fcc-editable-region--
    <input type="submit" value="Submit" />
  </form>
</body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
