---
id: 62ff8b9dab5ac88e4d3d43a3
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

按照可访问性最佳实践，使用 `for` 属性将 `input` 元素和 `label` 元素关联在一起。

使用 `first-name`、`last-name`、`email` 和 `new-password` 作为相应的 `id` 属性的值。

# --hints--

第一个 `input` 元素应该有一个 `first-name` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[0]?.matches('#first-name'))
```

第二个 `input` 元素应该有一个 `last-name` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[1]?.matches('#last-name'))
```

第三个 `input` 元素应该有一个 `email` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[2]?.matches('#email'))
```

第四个 `input` 元素应该有一个 `new-password` `id`。

```js
assert.isTrue(document.querySelectorAll('input')?.[3]?.matches('#new-password'))
```

第一个 `label` 元素应该有一个 `for` 属性，值为 `first-name`。

```js
assert.isTrue(document.querySelectorAll('label')?.[0]?.matches('label[for="first-name"]'))
```

第二个 `label` 元素应该有一个 `for` 属性，值为 `last-name`。

```js
assert.isTrue(document.querySelectorAll('label')?.[1]?.matches('label[for="last-name"]'))
```

第三个 `label` 元素应该有一个 `for` 属性，值为 `email`。

```js
assert.isTrue(document.querySelectorAll('label')?.[2]?.matches('label[for="email"]'))
```

第四个 `label` 元素应该有一个 `for` 属性，值为 `new-password`。

```js
assert.isTrue(document.querySelectorAll('label')?.[3]?.matches('label[for="new-password"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label>Enter Your First Name: <input /></label>
        <label>Enter Your Last Name: <input /></label>
        <label>Enter Your Email: <input /></label>
        <label>Create a New Password: <input /></label>
      </fieldset>
--fcc-editable-region--
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```

