---
id: 60eebd07ea685b0e777b5583
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

¡Bienvenido al proyecto de formulario de registro! Para iniciar, añada la declaración `!DOCTYPE html` en la parte superior del documento para que el navegador sepa qué tipo de documento está utilizando.

# --hints--

Su código debe contener la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Debes incluir un espacio despues de la referencia `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Debe definir el tipo de documento para que sea `html`.

```js
assert(code.match(/html/gi));
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert(code.match(/html\s*>/gi));
```

Su declaración `DOCTYPE` debe estar al comienzo de su HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
