---
id: 60f0286404aefb0562a4fdf9
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Añade un `title` y un elemento `meta` dentro del elemento `head`. Dale a tu proyecto el título de `Registration Form` y añade el atributo `charset` con el valor `utf-8` a tu elemento `meta`.

# --hints--

Tu código debe tener un elemento `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Tu elemento `title` debe estar dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Tu proyecto debe tener el título `Registration Form`. Recuerda que las mayúsculas, minúsculas y la ortografía son importantes para el título.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Debe crear un elemento `meta` dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Debes dar al elemento `meta` un atributo `charset` con el valor `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Tu elemento `meta` debe ser un elemento vacío, no tiene una etiqueta de cierre `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
