---
id: 60fab8367d35de04e5cb7929
title: Paso 31
challengeType: 0
dashedName: step-31
---

# --description--

Añadir el texto `I accept the terms and conditions` inmediatamente después del elemento `input` en el nuevo `label`. Luego, enlaza el texto `terms and conditions` a la siguiente ubicación:

```md
https://www.freecodecamp.org/news/terms-of-service/
```

# --hints--

Debes añadir el texto `I accept the terms and conditions` al `label` que sigue al tercer `fieldset`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label')?.innerText.trim(), 'I accept the terms and conditions');
```

Debes usar un elemento `a` para enlazar a los términos y condiciones.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label a'));
```

Debes colocar el nuevo texto inmediatamente después del elemento `input` en el `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input + a'));
```

Debes darle al elemento `a` un `href` de `https://www.freecodecamp.org/news/terms-of-service/`.

```js
assert.match(document.querySelector('fieldset:nth-child(3) + label > input + a')?.href, /https:\/\/www\.freecodecamp\.org\/news\/terms-of-service\/?/);
```

Debes envolver solo el elemento `a` alrededor del texto `terms and conditions`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent.trim(), 'terms and conditions');
```

El texto dentro de tu elemento ancla tiene espacios adicionales al inicio o al final. Los únicos espacios en el texto de anclaje deben estar entre las palabras `terms`, `and`, y `conditions`.

```js
const nestedAnchor = document.querySelector('fieldset:nth-child(3) + label > input + a')?.textContent;
const innerContent = nestedAnchor?.innerHTML;
assert.isNotTrue(/^\s+|\s+$/.test(innerContent));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--
      <label for="terms-and-conditions"><input id="terms-and-conditions" type="checkbox" required /></label>
--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
