---
id: 62cc5b1779e4d313466f73c5
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Anida un elemento `link` dentro del elemento `head`. Dale un atributo `rel` con el valor de `stylesheet` y un atributo `href` con un valor de `styles.css`.

# --hints--

Su código debe tener un elemento <code>link`.

```js
assert.isNotNull(document.querySelector('link'));
```

No debe cambiar sus etiquetas `head` existentes. Asegúrate de que no borras la etiqueta de cierre.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Debes tener un elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Tu elemento `link` debe ser un elemento vacío, no debe tener una etiqueta de cierre `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Tu elemento `link` debe estar dentro de tu elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Su elemento `link` debe tener un atributo `rel` con el valor `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Tu elemento `link` debe tener un atributo `href` con el valor `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
