---
id: 697fc395f3d85f14b05cd9ca
title: Rejesha Hadithi Inayolingana Kutoka Katika Safu ya Vipande vya Hadithi
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

Katika maabara hii, utarejesha hadithi inayolingana kutoka katika safu ya vipande vya hadithi vilivyoharibika.

Utajifunza kutumia mizunguko kwa kutekeleza algoriti za msingi za safu ya data kutoka mwanzo.

Utafanya kazi na safu za vitu vya vipande vya hadithi. Kila kitu cha kipande kina sifa zifuatazo:

| Sifa | Maelezo | Thamani ya Mfano |
|-|-|-|
| `id` | Nambari kamili chanya inayobainisha nafasi ya kipande katika hadithi | `3` |
| `text` | Maudhui halisi ya hadithi | `"and I use Arch btw.\""` |

Katika maabara hii, unapatiwa safu ya data iliyojazwa tayari iitwayo `shuffledFragments`.

**Lengo**: Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za Mtumizi**

1. Haufai kubadilisha safu ya data iliyojazwa tayari ya `shuffledFragments`.
2. Unapaswa kuunda kitendakazi kiitwacho `compactFragments` ambacho hupokea safu ya data ya vipande na kurudisha safu mpya ya data ambayo vipengele vyote vya undefined vimeondolewa. Ikiwa kitendakazi kitaondoa kipengele chochote cha undefined, kinapaswa kuandika ujumbe kwenye konsoli. Ujumbe unapaswa kuanza na kiambishi awali `[COMPACTED]`.
3. Unapaswa kutangaza kigezo kiitwacho `compactedShuffledFragments` na kukipa matokeo ya kuita `compactFragments` na safu ya data ya `shuffledFragments`.
4. Unapaswa kuunda kitendakazi kiitwacho `sortFragments` ambacho hupokea safu ya data ya vipande isiyo na vipengele vya `undefined` na kurudisha safu mpya ya data iliyopangwa kwa sifa ya `id` kwa mpangilio wa kupanda, huku vipande vyenye `id` sawa vikibaki katika mpangilio wake wa awali. Haupaswi kutumia njia ya `sort` iliyojengwa ndani ya JavaScript.
5. Unapaswa kutangaza kigezo kiitwacho `sortedFragments` na kukipa matokeo ya kuita `sortFragments` na safu ya data ya `compactedShuffledFragments`.
6. Unapaswa kuunda kitendakazi kiitwacho `dedupeFragments` ambacho hupokea safu ya data iliyopangwa ya vipande na kurudisha safu mpya ya data ambayo rudufu zimeondolewa, kikibakiza tu tokeo la kwanza. Unapaswa kufafanua rudufu kama vipande viwili au zaidi vyenye `id` sawa. Kwa kila `id` inayoondolewa rudufu, kitendakazi kinapaswa kuandika ujumbe kwenye konsoli. Ujumbe unapaswa kuanza na kiambishi awali `[DEDUPED]`.
7. Unapaswa kutangaza kigezo kiitwacho `dedupedFragments` na kukipa matokeo ya kuita `dedupeFragments` na safu ya data ya `sortedFragments`.
8. Unapaswa kuunda kitendakazi kiitwacho `fillMissingFragments` ambacho hupokea safu ya data iliyopangwa ya vipande na kurudisha safu mpya ya data ambayo vipande vilivyokosekana vimejazwa kwa vitu vya nafasi. Unapaswa kufafanua vipande vilivyokosekana kama mapengo katika mfululizo kati ya `id` ya chini kabisa na ya juu kabisa. Vitu vya nafasi vinapaswa kuwa na muundo wa `{ id: missingId, text: "[...]" }`. Kwa kila kitu cha nafasi kilichoongezwa, kitendakazi kinapaswa kuandika ujumbe kwenye konsoli. Ujumbe unapaswa kuanza na kiambishi awali `[FILLED]`.
9. Unapaswa kutangaza kigezo kiitwacho `filledFragments` na kukipa matokeo ya kuita `fillMissingFragments` na safu ya data ya `dedupedFragments`.
10. Unapaswa kuunda kitendakazi kiitwacho `assembleStory` ambacho hupokea safu ya data iliyopangwa ya vipande na kurudisha mfuatano mmoja wenye maandishi yote ya vipande, yakitenganishwa kwa mistari mipya.
11. Unapaswa kutumia `assembleStory` na `filledFragments` yako ili kuonyesha hadithi kamili kwenye konsoli.
12. Vitendakazi vyako `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments` na `assembleStory` havipaswi kubadilisha safu ya data ambayo vinaitwa nayo.

**Mfano**

Hapa kuna mfano wa safu ya data inayojumuisha vipande vya hadithi:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

Baada ya kurejesha hadithi kutoka `exampleArray`, itakuwa kama ifuatavyo:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

Haufai kubadilisha safu ya data iliyojazwa tayari ya `shuffledFragments`.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

Unapaswa kuwa na kitendakazi cha `compactFragments`.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` inapaswa kuandika ujumbe unaoanza na `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wakati kitendakazi cha `compactFragments` kinapoitwa na safu ya data yenye vipengele vya `undefined`, kinapaswa kuandika ujumbe kwenye konsoli unaoanza na `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` inapaswa kuandika ujumbe unaoanza na `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wakati kitendakazi cha `compactFragments` kinapoitwa na safu yenye nafasi tupu, kinapaswa kuandika ujumbe kwenye konsoli unaoanza na `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Wakati kitendakazi cha `compactFragments` kinapoitwa na safu `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`, hakipaswi kuandika ujumbe wowote kwenye konsoli unaoanza na `[COMPACTED]`.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` inapaswa kurudisha `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments` haipaswi kubadilisha safu ya data iliyotolewa.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments` inapaswa kuwa safu ya data isiyo na vipengele vya `undefined`.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

Unapaswa kuwa na kitendakazi cha `sortFragments`.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` inapaswa kurudisha `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments` haipaswi kutumia njia ya `.sort()` iliyojengwa ndani.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments` haipaswi kubadilisha safu ya data iliyotolewa.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments` inapaswa kuwa na kila kipengele chenye `id` kidogo au sawa na `id` ya kipengele kinachofuata.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

Unapaswa kuwa na kitendakazi cha `dedupeFragments`.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` inapaswa kurudisha `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` inapaswa kuandika ujumbe wa `[DEDUPED]` kwa kila `id` iliyorudufiwa.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` inapaswa kuandika ujumbe wa `[DEDUPED]` wakati `id` ya mwisho imerudufiwa.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

Wakati `dedupeFragments` inapoitwa na safu yenye rudufu, inapaswa kuandika ujumbe unaoanza na `[DEDUPED]` kwa kila `id` iliyorudufiwa.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments` haipaswi kubadilisha safu ya data iliyotolewa.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments` haipaswi kuwa na thamani za `id` rudufu.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

Unapaswa kuwa na kitendakazi cha `fillMissingFragments`.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` inapaswa kurudisha `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` inapaswa kuandika ujumbe unaoanza na `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

Wakati `fillMissingFragments` inaongeza kiwakilishi, inapaswa kuandika ujumbe unaoanza na `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments` haipaswi kubadilisha safu ya data iliyotolewa.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments` haipaswi kuwa na mapengo katika mfululizo wa `id`.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

Unapaswa kuwa na kitendakazi cha `assembleStory`.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` inapaswa kurudisha mfuatano wa herufi `"Hello\nWorld"`.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory` haipaswi kubadilisha safu ya data iliyotolewa.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

Unapaswa kuita `console.log` na matokeo ya `assembleStory(filledFragments)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
