---
id: 66f2836c459cfb16ae76f24f
title: 피라미드 생성기 만들기
challengeType: 26
dashedName: lab-pyramid-generator
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 세 개의 인수를 받는 `pyramid`라는 이름의 함수가 있어야 합니다.
1. 첫 번째 인수는 피라미드에서 반복할 패턴 문자를 나타내는 문자열이어야 합니다.
1. 두 번째 인수는 피라미드의 행 수를 나타내는 정수여야 합니다.
1. 세 번째 인수는 불리언 값이어야 합니다.
1. `pyramid` 함수는 세 번째 인수가 `false`일 때 꼭대기가 위를 향하는 피라미드를 만들도록 패턴 문자를 반복하고 배열한 문자열을 반환해야 합니다.
1. 세 번째 인수가 `true`일 때는 꼭대기가 아래를 향하는 피라미드를 만들어야 합니다.
1. 꼭대기 행은 패턴 문자가 한 개 있어야 하며, 그 외 각 행은 이전 행보다 패턴 문자가 두 개씩 더 있어야 합니다.
1. 각 행은 중앙 문자가 같은 열에 오도록 충분한 공백으로 시작해야 하며, 행 끝에는 공백이 없어야 합니다.
1. 피라미드는 줄바꿈 문자로 시작하고 끝나야 합니다.

예를 들어, `pyramid("o", 4, false)`를 호출하면 다음과 같은 출력이 나와야 합니다:

```js

   o
  ooo
 ooooo
ooooooo

```

# --hints--

`pyramid`라는 이름의 함수가 있어야 합니다.

```js
assert.isFunction(pyramid);
```

`pyramid` 함수는 세 개의 매개변수를 가져야 합니다.

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)`는 `"\n   o\n  ooo\n ooooo\nooooooo\n"`를 반환해야 합니다.

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)`는 `"\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n"`를 반환해야 합니다.

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
    const rows = []
    for (let i = 1; i <= count; i++) {
        if (isInverted) {
            rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
        } else {
            rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
        }        
    }
    return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```
