---
id: 646c47867800472a4ed5d2ea
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

從基本的 HTML 結構開始。為你的文件添加 `DOCTYPE` 的 `html` 引用，並新增一個 `html` 元素，其 `lang` 屬性設為 `en`。此外，在 `html` 元素內新增 `head` 和 `body` 元素。

# --hints--

你應該有 `html` 的 `DOCTYPE` 宣告。

```js
assert(code.match(/<!DOCTYPE html>/i));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開啟的 `html` 頁籤，且其 `lang` 屬性為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你應該有一個結尾的 `html` 頁籤。

```js
assert(code.match(/<\/html>/i));
```

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body>/i));
```

你的 `body` 元素應該位於 `head` 元素之後。

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

你的 `head` 和 `body` 元素應該位於 `html` 元素內。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
