---
id: 646c5ffef5598d449b52ec12
title: 第 19 步
challengeType: 0
dashedName: step-19
---

# --description--

現在你應該開始製作貓的耳朵。會有一個右耳和一個左耳，每個耳朵裡面都會有一個內耳。

在你的 `.cat-head` 元素內，創建一個帶有類別 `cat-ears` 的 `div` 元素。

# --hints--

你不應該更改具有 `cat-head` 類別的現有 `div` 元素。

```js
assert(document.querySelectorAll('div.cat-head').length === 1);
```

你應該在你的 `.cat-head` 元素內創建一個 `div` 元素。

```js
assert(document.querySelectorAll('.cat-head div').length === 1);
```

你的 `div` 元素應該有 `cat-ears` 類別。

```js
assert(document.querySelectorAll('.cat-head div')[0]?.getAttribute('class') === 'cat-ears');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      --fcc-editable-region--
      <div class="cat-head">
        
      </div>
      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
