---
id: 646c586be7180e40ddf74ff6
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

為你的 `body` 元素設定 `background-color` 為 `#c9d2fc`。

# --hints--

你應該有一個 `body` 選擇器。

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

你的 `body` 選擇器應該有一個 `background-color` 屬性設定為 `#c9d2fc`。

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.backgroundColor === 'rgb(201, 210, 252)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
* {
  box-sizing: border-box;
}

--fcc-editable-region--

--fcc-editable-region--
```

