---
id: 646ceb843412c74edee27a79
title: 第 22 步
challengeType: 0
dashedName: step-22
---

# --description--

現在你將學習一個用來繪製三角形的 CSS 技巧。

這將用於在貓咪上繪製耳朵。

使用類別選擇器，為 `.cat-right-ear` 元素設定 `height` 和 `width` 屬性為 `100px`。將 `background-color` 設定為白色。將左邊和右邊的邊框設定為 `35px solid blue`。將上邊和下邊的邊框設定為 `35px solid red`。

# --hints--

你應該有一個 `.cat-right-ear` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你的 `.cat-right-ear` 選擇器應該有一個 `height` 屬性設定為 `100px`。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height, '100px')
```

你的 `.cat-right-ear` 選擇器應該有一個 `width` 屬性設定為 `100px`。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width, '100px')
```

你的 `.cat-right-ear` 選擇器應該有一個 `background-color` 屬性設定為 `white`。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor, 'white')
```

你的 `.cat-right-ear` 選擇器應該有一個設定為 `35px solid blue` 的 `border-left` 屬性。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft, '35px solid blue')
```

你的 `.cat-right-ear` 選擇器應該有一個設定為 `35px solid blue` 的 `border-right` 屬性。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight, '35px solid blue')
```

你的 `.cat-right-ear` 選擇器應該有一個設定為 `35px solid red` 的 `border-top` 屬性。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop, '35px solid red')
```

你的 `.cat-right-ear` 選擇器應該有一個設定為 `35px solid red` 的 `border-bottom` 屬性。別忘了添加分號。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom, '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
