---
id: 646c48df8674cf2b91020ecb
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，為你的 `meta` 頁籤添加 `charset` 屬性為 `utf-8`。同時，為你的 `title` 元素添加文字 `fCC Cat Painting`。

# --hints--

你應該添加一個單一的 `meta` 元素。

```js
assert(document.querySelectorAll('meta').length === 1);
```

你的 `meta` 元素應該有一個設定為 `utf-8` 的 `charset` 屬性。

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

你應該添加一個單一的 `title` 元素。

```js
assert(document.querySelectorAll('title').length === 1);
```

你的 `title` 元素應該有文字 `fCC Cat Painting`。請注意拼寫和大小寫很重要。

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
