---
id: 6476fc5cf14b276e6d04e82a
title: 步驟 13
challengeType: 0
dashedName: step-13
---

# --description--

使用類別選擇器為你的 `.box` 元素設定寬度為 `200px`、高度為 `600px` 和背景顏色為 `#000`。同時，為它設定 `position` 為 `absolute`、`top` 為 `800px` 和 `left` 為 `650px`。 

# --hints--

你應該有一個 `.box` 選擇器。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box'))
```

你的 `.box` 選擇器應該有一個設定為 `200px` 的 `width` 屬性。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.width === '200px')
```

你的 `.box` 選擇器應該有一個設定為 `600px` 的 `height` 屬性。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.height === '600px')
```

你的 `.box` 選擇器應該設定 `background-color` 屬性為 `#000`。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.backgroundColor === 'rgb(0, 0, 0)')
```

你的 `.box` 選擇器應該有一個設定為 `absolute` 的 `position` 屬性。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.position === 'absolute')
```

你的 `.box` 選擇器應該有一個設定為 `650px` 的 `left` 屬性。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.left === '650px')
```

你的 `.box` 選擇器應該有一個設定為 `800px` 的 `top` 屬性。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.box')?.top === '800px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: relative;
  top: 100px;
  left: 100px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```

