---
id: 6476fb7cbaafa36d65e9cf35
title: 步驟 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` 是一個 `position` 屬性值，讓你可以將元素固定在頁面上，無論使用者如何捲動頁面。

你必須做更多標記才能了解 `fixed` 定位的運作方式。在你的 HTML 中，創建一個帶有類別 `box` 的 `div` 元素。

# --hints--

你應該有一個 `div` 元素。

```js
assert(document.querySelectorAll('div').length === 2);
```

你的 `div` 元素應該有 `box` 類別。

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
