---
id: 64770351e8586671ec0911f0
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

最後一個位置屬性值是 `sticky`。`sticky` 定位是 `relative` 和 `fixed` 定位的混合體。它允許元素根據捲動位置**黏附**在其包含元素或視窗內的特定位置。

將 `.cat-head` 的 `position` 屬性值改為 `sticky`，將 `top` 設為 `0`，然後移除 `left` 及其值。

**注意**：要查看 `sticky` 的運作方式，你必須在你的 `.cat-head` `div` 元素前後放置幾段文字。之後如果你向下捲動，就會看到 `.cat-head` 被固定在頂部並保持在那裡。

# --hints--

你的 `.cat-head` 選擇器應該有一個設定為 `sticky` 的 `position` 屬性。請確保你添加分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'sticky')
```

你的 `.cat-head` 選擇器應該有一個設定為 `0` 的 `top` 屬性。請確保你添加分號。

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '0px')
```

你的程式碼中不應該有 `left` 屬性及其值。

```js
assert.notMatch(code, /left:\s*100px;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: fixed;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
