---
id: 647d855ac12cd436059acd39
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

使用通用選擇器為你的 CSS 添加 `box-sizing: border-box;`。這確保元素在指定的寬度和高度中包含內距和邊框。

# --hints--

你應該有通用選擇器（`*`）。

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

你的通用選擇器應該有一個設定為 `border-box` 的 `box-sizing` 屬性。

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
