---
id: 6676a8a8d9861319ee901a90
title: 第 24 步
challengeType: 0
dashedName: step-24
---

# --description--

現在你可以開始製作你貓咪的耳朵。

透過移除 `.cat-right-ear` 選擇器及其所有屬性來清理你的審查程式碼。

使用類別選擇器，為 `.cat-left-ear` 元素的左右邊框各設定 `35px solid transparent`。同時，將底部邊框設定為 `70px solid #5e5e5e`。

# --hints--


你的 `.cat-right-ear` 選擇器不應該有 `height` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height)
```

你的 `.cat-right-ear` 選擇器不應該有 `width` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width)
```


你的 `.cat-right-ear` 選擇器不應該有 `border-left` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft)
```

你的 `.cat-right-ear` 選擇器不應該有 `border-right` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight)
```

你的 `.cat-right-ear` 選擇器不應該有 `border-top` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop)
```

你的 `.cat-right-ear` 選擇器不應該有 `border-bottom` 屬性。

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom)
```

你不應該有 `.cat-right-ear` 選擇器。

```js
assert(!new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

你應該有一個 `.cat-left-ear` 選擇器。

```js 
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear'))
```

你的 `.cat-left-ear` 選擇器應該有一個設定為 `35px solid transparent` 的 `border-left` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderLeft === '35px solid transparent')
```

你的 `.cat-left-ear` 選擇器應該有一個設定為 `35px solid transparent` 的 `border-right` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderRight === '35px solid transparent')
```

你的 `.cat-left-ear` 選擇器應該有一個設定為 `70px solid #5e5e5e` 的 `border-bottom` 屬性。

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-left-ear')?.borderBottom === '70px solid rgb(94, 94, 94)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 0;
  width: 0;
  border-left: 35px solid transparent;
  border-right: 35px solid transparent;
  border-top: 35px solid transparent;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
