---
id: 66cc14ca14e65e7a9ebecb80
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Da Ihr Taschenrechner jetzt zwei Zahlen addieren kann, ist es Zeit, eine Funktion zu erstellen, die zwei Zahlen subtrahiert.

Deklarieren Sie eine Funktion namens `calculateDifference`, die zwei Parameter, `num1` und `num2`, entgegennimmt.

Geben Sie innerhalb der Funktion die Differenz der beiden Parameter zurück.

# --hints--

Sie sollten eine Funktion namens `calculateDifference` haben.

```js
assert.isFunction(calculateDifference);
```

Ihre `calculateDifference`-Funktion sollte zwei Parameter namens `num1` und `num2` haben.

```js
const regex = __helpers.functionRegex('calculateDifference', ['num1', 'num2']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ihre `calculateDifference`-Funktion sollte die Differenz der beiden Parameter zurückgeben.

```js
assert.equal(calculateDifference(5, 2), 3);
assert.equal(calculateDifference(10, 5), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

--fcc-editable-region--

--fcc-editable-region--
```
