---
id: 66cc10958d72e57907f898b2
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Innerhalb Ihrer `calculateSum`-Funktion müssen Sie die Summe der beiden Parameter zurückgeben.

# --hints--

Ihre `calculateSum`-Funktion sollte die Summe der beiden Parameter zurückgeben.

```js
assert.equal(calculateSum(1, 2), 3);
assert.equal(calculateSum(5, 5), 10);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
