---
id: 66cc0a9e06e00b75d6782be9
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man eine Funktion aufruft (invoked). Eine Funktion aufzurufen bedeutet, den Code innerhalb der Funktion auszuführen.

Hier eine Erinnerung daran, wie man eine Funktion aufruft:

```js
function myFunction() {
  return "Hello World";
}

// function call
myFunction();
```

Fügen Sie eine `console.log`-Anweisung hinzu, und rufen Sie darin die `addTwoAndSeven`-Funktion auf.

Sie sollten nun die Summe von `2` und `7` in der Konsole sehen.

# --hints--

Sie sollten die `addTwoAndSeven`-Funktion innerhalb einer `console.log` aufrufen.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(addTwoAndSeven())";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

--fcc-editable-region--

--fcc-editable-region--
```
