---
id: 66cc12fa504b0479dac479a0
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Der Vorteil von Funktionen ist, dass sie mit verschiedenen Argumenten aufgerufen werden können, sodass Sie denselben Code mit unterschiedlichen Werten wiederverwenden können.

Hier sind einige Beispielaufrufe von Funktionen mit verschiedenen String-Argumenten:

```js
// function definition
function greetUser(name) {
  console.log(`Hello ${name}!`);
}

// function calls
greetUser("John"); // "Hello John!"
greetUser("Jane"); // "Hello Jane!"
```

Fügen Sie einen `console.log` hinzu, der die `calculateSum`-Funktion mit den Argumenten `2` und `5` aufruft.

# --hints--

Sie sollten einen `console.log` haben, der die `calculateSum`-Funktion mit den Argumenten `2` und `5` aufruft.

```js
assert.match(code, /console\.log\(calculateSum\(\s*2\s*,\s*5\)\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
    return num1 + num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
