---
id: 66cc21d23238dc8240a8a182
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Es wäre schön, wenn Ihr Taschenrechner das Quadrat einer Zahl berechnen könnte. Das Quadrat einer Zahl ist die Zahl multipliziert mit sich selbst.

Um das Quadrat einer Zahl in JavaScript zu berechnen, können Sie die <dfn>Math.pow()</dfn>-Methode verwenden. Die `Math.pow()`-Methode nimmt zwei Argumente: die Basiszahl und den Exponenten.

Sie können auch den Exponentationsoperator (`**`) verwenden, um das Quadrat einer Zahl zu berechnen.

Hier ist ein Beispiel:

```js
// base number is 5 
// the exponent is 2
Math.pow(5, 2); // 25

// using the exponentiation operator
5 ** 2; // 25
```

Deklarieren Sie eine Funktion namens `calculateSquare`, die einen `num`-Parameter entgegennimmt und das Quadrat von `num` zurückgibt.

# --hints--

Sie sollten eine Funktion namens `calculateSquare` haben.

```js
assert.isFunction(calculateSquare);
```

Ihre `calculateSquare`-Funktion sollte einen `num`-Parameter annehmen.

```js
const regex = __helpers.functionRegex('calculateSquare', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ihre `calculateSquare`-Funktion sollte das Quadrat von `num` zurückgeben.

```js
assert.equal(calculateSquare(5), 25);
assert.equal(calculateSquare(10), 100);
assert.equal(calculateSquare(3), 9);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

--fcc-editable-region--

--fcc-editable-region--
```
