---
id: 66cc0de2920bec775f610424
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Obwohl die Funktionen wie erwartet funktionieren, gibt es in Ihrem Code viele Wiederholungen.

Da Sie einen Taschenrechner bauen, möchten Sie nicht für jede mögliche Kombination von Zahlen, die Sie addieren möchten, eine eigene Funktion erstellen.

Es wäre besser, eine einzige wiederverwendbare Funktion zu erstellen, die zwei beliebige Zahlen zusammenaddieren kann.

Entfernen Sie den gesamten Code, den Sie bisher geschrieben haben.

Im nächsten Schritt werden Sie lernen, wie man mit Parametern und Argumenten in Funktionen arbeitet.

# --hints--

Sie sollten keine Funktion namens `addTwoAndSeven` haben.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

Sie sollten keinen Funktionsaufruf zu `addTwoAndSeven` haben.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

Sie sollten keine Funktion namens `addThreeAndFour` haben.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

Sie sollten keinen Funktionsaufruf zu `addThreeAndFour` haben.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
