---
id: 66cc06012e46aa6bc9b8c001
title: Schritt 1
challengeType: 1
dashedName: step-1
---

# --description--

In diesem Workshop überprüfen Sie, wie Sie mit Funktionen arbeiten, indem Sie eine Taschenrechner-App erstellen.

In vorherigen Lektionen haben Sie gelernt, Funktionen so zu deklarieren:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Beginnen Sie damit, eine Funktion namens `addTwoAndSeven` zu erstellen.

Sie können wählen, ob Sie die reguläre Funktionssyntax oder die Syntax der Pfeilfunktion verwenden möchten.

# --hints--

Sie sollten eine Funktion namens `addTwoAndSeven` deklarieren.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
