---
id: 66cc172af7159a7c67804544
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Jetzt ist es Zeit, die Multiplikationsfunktionalität zu Ihrem Taschenrechner hinzuzufügen.

Deklarieren Sie eine Funktion namens `calculateProduct`, die zwei Parameter `num1` und `num2` entgegennimmt.

Geben Sie innerhalb der Funktion das Produkt der beiden Parameter zurück.

Fügen Sie dann einen `console.log` hinzu, der die Funktion `calculateProduct` mit den Argumenten `13` und `5` aufruft.

# --hints--

Sie sollten eine Funktion namens `calculateProduct` deklarieren.

```js
assert.isFunction(calculateProduct);
```

Ihre `calculateProduct`-Funktion sollte zwei Parameter namens `num1` und `num2` haben.

```js
const explorer = await __helpers.Explorer(code);
const { calculateProduct } = explorer.allFunctions;
assert.lengthOf(calculateProduct?.parameters, 2);
assert.equal(calculateProduct?.parameters[0].toString(), "num1");
assert.equal(calculateProduct?.parameters[1].toString(), "num2");
```

Ihre `calculateProduct`-Funktion sollte das Produkt der beiden Parameter zurückgeben.

```js
assert.equal(calculateProduct(5, 2), 10);
assert.equal(calculateProduct(10, 5), 50);
```

Sie sollten einen `console.log` haben, der die Funktion `calculateProduct` mit den Argumenten `13` und `5` aufruft.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const validStatements = ["console.log(calculateProduct(13,5))", "console.log(calculateProduct(5,13))"];
assert.isTrue(validStatements.some(stmt => strippedCode.includes(stmt)));
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

--fcc-editable-region--

--fcc-editable-region--
```
