---
id: 66cc26907ff6908402af0149
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Um Ihre `calculateSquare`-Funktion zu testen, müssen Sie die Funktion ein paar Mal aufrufen.

Fügen Sie einen `console.log` hinzu, der die `calculateSquare`-Funktion mit dem Argument `2` aufruft.

Fügen Sie dann einen weiteren `console.log` hinzu, der die `calculateSquare`-Funktion mit dem Argument `9` aufruft.

# --hints--

Sie sollten einen `console.log` haben, der die `calculateSquare`-Funktion mit dem Argument `2` aufruft.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*2\s*\)\s*\)/);
```

Sie sollten einen `console.log` haben, der die `calculateSquare`-Funktion mit dem Argument `9` aufruft.

```js
assert.match(code, /console\.log\s*\(\s*calculateSquare\s*\(\s*9\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

--fcc-editable-region--

--fcc-editable-region--
```
