---
id: 66cc0f1ae40802781b2ea972
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

In den vorherigen Lektionen haben Sie gelernt, wie man mit Funktionsparametern arbeitet.

Ein Funktionsparameter ist eine Variable, die in der Deklaration der Funktion definiert wird und als Platzhalter dient.

Hier ist ein Beispiel für eine Funktion, die einen Parameter hat:

```js
// The parameter is `name`
function greetUser(name) {
  return `Hello, ${name}!`;
}
```

Deklarieren Sie eine Funktion namens `calculateSum`, die zwei Parameter, `num1` und `num2`, entgegennimmt.

# --hints--

Sie sollten eine Funktion namens `calculateSum` deklarieren.

```js
assert.isFunction(calculateSum);
```

Ihre `calculateSum`-Funktion sollte zwei Parameter, `num1` und `num2`, entgegennehmen.

```js
const explorer = await __helpers.Explorer(code);
const { calculateSum } = explorer.allFunctions;
assert.lengthOf(calculateSum?.parameters, 2);
assert.equal(calculateSum?.parameters[0].toString(), "num1");
assert.equal(calculateSum?.parameters[1].toString(), "num2");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
