---
id: 66cc2739f687e484d50bb6f1
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

Die letzte Funktionalität für Ihren Taschenrechner wird die Quadratwurzelfunktion sein.

Ein Quadrat ist eine Zahl, die mit sich selbst multipliziert wird. Die Quadratwurzel ist also die Zahl, die, wenn sie mit sich selbst multipliziert wird, die ursprüngliche Zahl ergibt.

Zum Beispiel ist die Quadratwurzel von `9` `3`, weil `3 * 3 = 9`.

Um die Quadratwurzel einer Zahl in JavaScript zu erhalten, können Sie die <dfn>Math.sqrt()</dfn>-Methode verwenden.

Hier ist ein Beispiel:

```js
// result: 3
Math.sqrt(9);
```

Deklarieren Sie eine Funktion namens `calculateSquareRoot`, die einen `num`-Parameter entgegennimmt und die Quadratwurzel von `num` zurückgibt.

# --hints--

Sie sollten eine Funktion namens `calculateSquareRoot` haben.

```js
assert.isFunction(calculateSquareRoot);
```

Ihre `calculateSquareRoot`-Funktion sollte einen `num`-Parameter annehmen.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Ihre `calculateSquareRoot`-Funktion sollte die Quadratwurzel von `num` zurückgeben.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
