---
id: 66cc16abc30b367ba1d1f9ef
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Jetzt ist es Zeit, Ihre `calculateDifference`-Funktion zu testen.

Beginnen Sie damit, eine `console.log` hinzuzufügen, die die `calculateDifference`-Funktion mit den Argumenten `22` und `5` aufruft.

Fügen Sie dann eine zweite `console.log` hinzu, die die `calculateDifference`-Funktion mit den Argumenten `12` und `1` aufruft.

Schließlich fügen Sie eine dritte `console.log` hinzu, die die `calculateDifference`-Funktion mit den Argumenten `17` und `9` aufruft.

# --hints--

Sie sollten eine `console.log` haben, die die `calculateDifference`-Funktion mit den Argumenten `22` und `5` aufruft.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*22\s*,\s*5\s*\)\s*\)\s*;?/);
```

Sie sollten eine `console.log` haben, die die `calculateDifference`-Funktion mit den Argumenten `12` und `1` aufruft.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*12\s*,\s*1\s*\)\s*\)\s*;?/);
```

Sie sollten eine `console.log` haben, die die `calculateDifference`-Funktion mit den Argumenten `17` und `9` aufruft.

```js
assert.match(code, /console\.log\s*\(\s*calculateDifference\s*\(\s*17\s*,\s*9\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

--fcc-editable-region--

--fcc-editable-region--
```
