---
id: 66cc1a3a39aef47d6473cb2f
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Der nächste Schritt ist, die Division in Ihren Taschenrechner einzubauen.

Deklarieren Sie eine Funktion namens `calculateQuotient`, die zwei Parameter `num1` und `num2` entgegennimmt.

Geben Sie innerhalb der Funktion den Quotienten der beiden Parameter zurück.

Fügen Sie dann einen `console.log` hinzu, der die Funktion `calculateQuotient` mit den Argumenten `7` und `11` aufruft.

# --hints--

Sie sollten eine Funktion namens `calculateQuotient` deklarieren.

```js
assert.isFunction(calculateQuotient);
```

Ihre Funktion `calculateQuotient` sollte zwei Parameter namens `num1` und `num2` haben.

```js
const explorer = await __helpers.Explorer(code);
const { calculateQuotient } = explorer.allFunctions;
assert.lengthOf(calculateQuotient?.parameters, 2);
assert.equal(calculateQuotient?.parameters[0].toString(), "num1");
assert.equal(calculateQuotient?.parameters[1].toString(), "num2");
```

Ihre Funktion `calculateQuotient` sollte den Quotienten der beiden Parameter zurückgeben.

```js
assert.equal(calculateQuotient(10, 2), 5);
assert.equal(calculateQuotient(100, 10), 10);
```

Sie sollten einen `console.log` haben, der die Funktion `calculateQuotient` mit den Argumenten `7` und `11` aufruft.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateQuotient(7,11))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));

--fcc-editable-region--

--fcc-editable-region--
```
