---
id: 66cc0ba5881acb7692cfc4de
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Jetzt ist es Zeit, eine weitere Funktion hinzuzufügen.

Deklarieren Sie eine Funktion namens `addThreeAndFour`, die die Summe von `3` und `4` zurückgibt.

Rufen Sie dann die `addThreeAndFour`-Funktion innerhalb eines `console.log` auf, um das Ergebnis zu sehen.

# --hints--

Sie sollten eine Funktion namens `addThreeAndFour` deklarieren.

```js
assert.isFunction(addThreeAndFour);
```

Ihre `addThreeAndFour`-Funktion sollte die Summe von `3` und `4` zurückgeben.

```js
assert.equal(addThreeAndFour(), 7);
```

Ihre `addThreeAndFour`-Funktion sollte die Summe mit dem `+`-Operator zurückgeben.

```js
const declarationStyle = /return\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const arrowStyle = /=>\s*\(?\s*(3\s*\+\s*4|4\s*\+\s*3)\s*\)?\s*;?/;
const combinedRegex = new RegExp(`(?:${declarationStyle.source})|(?:${arrowStyle.source})`);
assert.match(code, combinedRegex);
```

Sie sollten die `addThreeAndFour`-Funktion innerhalb eines `console.log` aufrufen.

```js
assert.match(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

--fcc-editable-region--

--fcc-editable-region--
```
