---
id: 66cc095dfe1682753d2ab030
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Eines der Konzepte, die Sie in vorherigen Lektionen gelernt haben, war, wie man Werte aus einer Funktion zurückgibt.

Hier eine Erinnerung daran, wie man einen Wert aus einer Funktion zurückgibt:

```js
function myFunction() {
  return "Hello World";
}
```

Geben Sie in Ihrer `addTwoAndSeven`-Funktion die Summe von `2` und `7` zurück.

# --hints--

Ihre `addTwoAndSeven`-Funktion sollte die Summe von `2` und `7` zurückgeben.

```js
assert.strictEqual(addTwoAndSeven(), 9);
```

Ihre `addTwoAndSeven`-Funktion sollte die Summe mit dem `+`-Operator zurückgeben.

```js
assert.match(code, /return\s*(2\s*\+\s*7|7\s*\+\s*2)\s*;?/);
```

# --seed--

## --seed-contents--

```js
function addTwoAndSeven() {  
--fcc-editable-region--

--fcc-editable-region--
}
```
