---
id: 66cdf76685e4cb5a8726e27b
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Um Ihre `calculateSum`-Funktion weiter zu testen, müssen Sie sie mit verschiedenen Argumenten aufrufen.

Fügen Sie eine weitere `console.log` hinzu, die die `calculateSum`-Funktion mit den Argumenten `10` und `10` aufruft.

Fügen Sie unter dieser `console.log` eine weitere `console.log` hinzu, die die `calculateSum`-Funktion mit den Argumenten `5` und `5` aufruft.

# --hints--

Sie sollten eine `console.log` haben, die die `calculateSum`-Funktion mit den Argumenten `10` und `10` aufruft.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(10,10))";
assert.include(strippedCode, expectedStatement);
```

Sie sollten eine `console.log` haben, die die `calculateSum`-Funktion mit den Argumenten `5` und `5` aufruft.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(calculateSum(5,5))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
--fcc-editable-region--

--fcc-editable-region--
```
