---
id: 66bb6a9c2dd58b73cd759034
title: 建置一個成績簿應用程式
challengeType: 26
dashedName: build-a-gradebook-app
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該有一個名為 `getAverage` 的函式（程式），它接受一個測驗分數的陣列作為參數並傳回平均分數。
1. 你應該有一個名為 `getGrade` 的函式（程式），它接受學生分數作為參數，並傳回一個字串，表述根據分數的等第。以下是分數與其對應的等第：分數區間 等第 `100` `"A+"` `90 - 99` `"A"` `80 - 89` `"B"` `70 - 79` `"C"` `60 - 69` `"D"` `0 - 59` `"F"`

1. 你應該有一個名為 `hasPassingGrade` 的函式（程式），它接受一個分數作為參數，並根據該分數是否對應及格成績傳回 `true` 或 `false`。
1. `hasPassingGrade` 函式（程式）應該使用 `getGrade` 函式（程式）來取得字母成績，並用它來判斷成績是否及格。及格成績是任何不等於 `"F"` 的成績。
1. You should have a function named `studentMsg` that takes an array of scores and a student score as the parameters. The function should return a string with the format:

    - 「類別平均：average-goes-here。你的成績：grade-goes-here。你通過了課程。」
    - 「類別平均：average-goes-here。你的成績：grade-goes-here。你未通過此課程。」如果學生未通過此課程。

    Replace `average-goes-here` with the average of total scores and `grade-goes-here` with the student's grade.
    Use `getAverage` to get the average score and `getGrade` to get the student's grade.

# --hints--

你應該有一個名為 `getAverage` 的函式（程式）。

```js
assert.isFunction(getAverage);
```

你的 `getAverage` 函式應該傳回一個數字。

```js
assert.isNumber(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` 應該傳回 `71.7`。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` 應該傳回 `85.4`。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` 應該傳回 `92.4`。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage([10, 20, 30, 40, 55, 65, 75, 83])` 應該傳回 `47.25`。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 55, 65, 75, 83]), 47.25);
```

`getAverage([10, 20, 30, 40, 50, 60, 70, 97])` 應該傳回 `47.125`。

```js
assert.strictEqual(getAverage([10, 20, 30, 40, 50, 60, 70, 97]), 47.125);
```

你的 `getAverage` 函式（程式）應該傳回測驗分數的平均分數。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

你應該有一個名為 `getGrade` 的函式（程式）。

```js
assert.isFunction(getGrade);
```

你的 `getGrade` 函式應該在分數為 `100` 時傳回 `"A+"`。

```js
assert.strictEqual(getGrade(100), "A+");
```

如果分數介於 `90` 和 `99` 之間，你的 `getGrade` 函式（程式）應該傳回 `"A"`。

```js
assert.strictEqual(getGrade(90), "A");
assert.strictEqual(getGrade(94), "A");
assert.strictEqual(getGrade(99), "A");
```

如果分數介於 `80` 和 `89` 之間，你的 `getGrade` 函式（程式）應該傳回 `"B"`。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

如果分數介於 `70` 和 `79` 之間，你的 `getGrade` 函式應該傳回 `"C"`。

```js
assert.strictEqual(getGrade(70), "C");
assert.strictEqual(getGrade(75), "C");
assert.strictEqual(getGrade(79), "C");
```

如果分數介於 `60` 和 `69` 之間，你的 `getGrade` 函式（程式）應該傳回 `"D"`。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

如果分數介於 `0` 和 `59` 之間，你的 `getGrade` 函式（程式）應該傳回 `"F"`。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
assert.strictEqual(getGrade(59), "F");
```


你應該有一個名為 `hasPassingGrade` 的函式（程式）。

```js
assert.isFunction(hasPassingGrade);
```

你的 `hasPassingGrade` 函式應該傳回一個布林值。

```js
assert.isBoolean(hasPassingGrade(100));
```

你的 `hasPassingGrade` 函式（程式）如果成績是 `"A"`，應該傳回 `true`。

```js
assert.isTrue(hasPassingGrade(100));
```

如果成績是 `"F"`，你的 `hasPassingGrade` 函式應該傳回 `false`。

```js
assert.isFalse(hasPassingGrade(53));
```

你的 `hasPassingGrade` 函式（程式）應該為所有及格成績傳回 `true`。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

你應該有一個名為 `studentMsg` 的函式（程式）。

```js
assert.isFunction(studentMsg);
```

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` 應該傳回以下訊息：「類別平均：71.7。你的成績：F。你未通過課程。」。

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` 應該傳回以下訊息：「Class average: 50.8. Your grade: A+. You passed the course.」。

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A+. You passed the course.");
```

`studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85)` 應該傳回以下訊息：「類別平均：48.25。你的成績：B。你通過了課程。」。

```js
assert.strictEqual(studentMsg([12, 22, 32, 42, 52, 62, 72, 92], 85), "Class average: 48.25. Your grade: B. You passed the course.");
```

`studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75)` 應該傳回以下訊息：「Class average: 45.625. Your grade: C. You passed the course.」。

```js
assert.strictEqual(studentMsg([15, 25, 35, 45, 55, 60, 70, 60], 75), "Class average: 45.625. Your grade: C. You passed the course.");
```

你的 `studentMsg` 函式應該根據學生的分數和類別平均值傳回正確的訊息。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function getAverage(scores) {
    let sum = 0;

    for (const score of scores) {
      sum += score;
    }

    return sum / scores.length;
  }

  function getGrade(score) {
    if (score === 100) {
      return "A+";
    } else if (score >= 90) {
      return "A";
    } else if (score >= 80) {
      return "B";
    } else if (score >= 70) {
      return "C";
    } else if (score >= 60) {
      return "D";
    } else {
      return "F";
    }
  }

  function hasPassingGrade(score) {
    return getGrade(score) !== "F";
  }

  function studentMsg(totalScores, studentScore) {
    let average = getAverage(totalScores);
    let grade = getGrade(studentScore);

    return `Class average: ${average}. Your grade: ${grade}. You ${
      hasPassingGrade(studentScore) ? "passed" : "failed"
    } the course.`;
  }
```
