---
id: 68cabc534d2e741f33a3fe1d
title: Jenga kifuatiliaji cha takwimu za mhusika wa mchezo
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

Katika maabara hii, utafanya kifuatiliaji cha takwimu za mhusika wa mchezo. Programu itakuwezesha kuunda mhusika mwenye sifa maalum, kusahihisha sifa hizo, na kupata takwimu za sasa za mhusika.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite mitihani yote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unda darasa la vitu liitwalo `GameCharacter` linalowakilisha mhusika wa mchezo na kusimamia takwimu za mhusika.
2. Wakati darasa hilo linapotengenezwa, kitu kipya cha `GameCharacter` kinapaswa kuwa na sifa zifuatazo:
    - `_name` iwe na mfuatano wa herufi uliotolewa wakati wa kuanzisha.
    - `_health` iwe na `100`.
    - `_mana` iwe na `50`.
    - `_level` iwe na `1`.
3. Unda sifa ya `name` kwa upatikanaji wa kusoma tu wa jina la mhusika.
4. Kwa sifa ya `health`:
     - Tafsiri getter itoe afya ya sasa.
     - Tafsiri setter isiruhusu afya kuwekwa chini ya `0`, na iweze kufikia kiwango cha juu cha `100`.
5. Kwa sifa ya `mana`:
     - Tafsiri getter itoe mana ya sasa.
     - Tafsiri setter isiruhusu mana kuwekwa chini ya `0`, na iweze kufikia kiwango cha juu cha `50`.
6. Unda getter kwa `level` kurudisha kiwango cha sasa cha mhusika.
7. Tafsiri njia (method) liitwalo `level_up` ambayo:
    - Inaongeza kiwango cha mhusika kwa 1.
    - Inarejesha afya kuwa `100` na mana kuwa `50` kwa kutumia setters zao.
    - Inachapisha ujumbe kwa muundo wa `<name> leveled up to <level>!` (ambapo `<name>` na `<level>` zinapaswa kubadilishwa na jina la mhusika na kiwango kipya, mtawalia).
8. Tafsiri njia ya `__str__` inayorudisha mfuatano wa herufi uliopangwa unaojumuisha:
    - Jina la mhusika.
    - Kiwango cha mhusika.
    - Afya ya sasa ya mhusika.
    - Mana ya sasa ya mhusika.
   Kwa mfano, mhusika aliyeitwa `Kratos`, mara tu baada ya kuanzishwa, anapaswa kuwakilishwa kama ifuatavyo:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Mfano wa matumizi

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Unapaswa kuwa na darasa la vitu la `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Unapaswa kuwa na njia ya `__init__` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

Njia ya `__init__` inapaswa kuwa na vigezo viwili, ambapo cha kwanza ni `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Unapaswa kuweka `self._name` kuwa mfuatano wa herufi uliotolewa wakati wa kuanzisha ndani ya njia ya `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Unapaswa kuweka `100` kuwa `self._health` ndani ya njia ya `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Unapaswa kuweka `50` kuwa `self._mana` ndani ya njia ya `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Unapaswa kuweka `1` kuwa `self._level` ndani ya njia ya `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Unapaswa kuwa na njia ya `name` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

Njia ya `name` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Unapaswa kurudisha `self._name` katika njia ya `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

Njia ya `name` inapaswa kuwa na dekoreta ya `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Unapaswa kuwa na njia ya `health` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

Njia ya `health` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Unapaswa kurudisha `self._health` katika njia ya `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

Njia ya `health` inapaswa kuwa na dekoreta ya `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Unapaswa kuunda njia ya setter ya `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

Njia ya setter ya `health` inapaswa kuwa na dekoreta ya `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

Njia ya setter ya `health` inapaswa kuwa na vigezo viwili, ambapo cha kwanza ni `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Unapaswa kuweka `self._health` kuwa `0` ikiwa thamani iliyotolewa kwa setter ya `health` ni chini ya `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Unapaswa kuweka `self._health` kuwa `100` ikiwa thamani iliyotolewa kwa setter ya `health` ni zaidi ya `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Unapaswa kuweka `self._health` kuwa thamani iliyotolewa kwa setter ya `health` wakati thamani iko kati ya `0` na `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
    char.health = 0
    assert char.health == 0
    char.health = 100
    assert char.health == 100
`)})
```

Unapaswa kuwa na njia ya `mana` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

Njia ya `mana` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Unapaswa kurudisha `self._mana` katika njia ya `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

Njia ya `mana` inapaswa kuwa na dekoreta ya `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Unapaswa kuunda njia ya setter ya `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

Njia ya setter ya `mana` inapaswa kuwa na dekoreta ya `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

Njia ya setter ya `mana` inapaswa kuwa na vigezo viwili, ambapo cha kwanza ni `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Unapaswa kuweka `self._mana` kuwa `0` ikiwa thamani iliyotolewa kwa setter ya `mana` ni chini ya `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Unapaswa kuweka `self._mana` kuwa `50` ikiwa thamani iliyotolewa kwa setter ya `mana` ni zaidi ya `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Unapaswa kuweka `self._mana` kuwa thamani iliyotolewa kwa setter ya `mana` ikiwa thamani iko kati ya `0` na `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Unapaswa kuwa na njia ya `level` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

Njia ya `level` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Unapaswa kurudisha `self._level` katika njia ya `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

Njia ya `level` inapaswa kuwa na dekoreta ya `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Unapaswa kuwa na njia ya `level_up` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

Njia ya `level_up` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Unapaswa kuongeza `self._level` kwa `1` katika njia ya `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Unapaswa kuweka `self.health` kuwa `100` katika njia ya `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Unapaswa kuweka `self.mana` kuwa `50` katika njia ya `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

Njia ya `level_up` inapaswa kuchapisha `<name> leveled up to <level>!` (ambapo `<name>` na `<level>` zinapaswa kubadilishwa na jina la mhusika na kiwango kipya, mtawalia).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Unapaswa kuwa na njia ya `__str__` katika darasa la `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

Njia ya `__str__` inapaswa kuwa na kigezo kimoja tu, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Njia yako ya `__str__` inapaswa kurudisha mfuatano wa herufi wenye takwimu za mhusika kwa kutumia muundo uliotolewa.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
