---
id: 66f17db06803d11a1bd19a20
title: Crea un gioco quiz
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Devi creare un array chiamato `questions`.
1. L'array `questions` deve contenere almeno cinque oggetti, ognuno con le chiavi `category`, `question`, `choices` e `answer`.
1. La chiave `category` deve avere il valore di una stringa che rappresenta una categoria di domanda.
1. La chiave `question` deve avere il valore di una stringa che rappresenta una domanda.
1. La chiave `choices` deve avere il valore di un array contenente tre stringhe, che sono risposte alternative alla domanda.
1. La chiave `answer` deve avere il valore di una stringa, che rappresenta la risposta corretta alla domanda. Inoltre, il valore di `answer` deve essere incluso nell'array `choices`.
1. Devi avere una funzione chiamata `getRandomQuestion` che prende un array di domande come parametro e restituisce un oggetto domanda casuale dall'array.
1. Devi avere una funzione chiamata `getRandomComputerChoice` che prende l'array delle scelte disponibili come parametro e restituisce una risposta casuale alla domanda selezionata.
1. Devi avere una funzione chiamata `getResults` che prende l'oggetto domanda come primo parametro e la scelta del computer come secondo parametro. La funzione deve restituire `The computer's choice is correct!` se la risposta è corretta. Altrimenti, restituisce `The computer's choice is wrong. The correct answer is: <correct-answer>`, dove `<correct-answer>` è il valore della risposta corretta alla domanda scelta.

# --hints--

Devi creare un array chiamato `questions`.

```js
assert.isArray(questions);
```

L'array `questions` deve contenere almeno cinque oggetti, ognuno con le chiavi `category`, `question`, `choices` e `answer`.

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

La chiave `category` deve avere il valore di una stringa che rappresenta una categoria di domanda.

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

La chiave `question` deve avere il valore di una stringa che rappresenta una domanda.

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

La chiave `choices` deve avere il valore di un array contenente tre stringhe diverse tra loro.

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

La chiave `answer` deve avere il valore di una stringa.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

Il valore di `answer` deve essere incluso nell'array `choices`.

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

Devi avere una funzione chiamata `getRandomQuestion` che prende un array di domande come parametro e restituisce un oggetto domanda casuale dall'array.

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

Devi avere una funzione chiamata `getRandomComputerChoice` che prende l'array delle scelte disponibili come parametro e restituisce una risposta casuale alla domanda selezionata.

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

Devi avere una funzione chiamata `getResults`.

```js
assert.isFunction(getResults);
```

La tua funzione `getResults` deve prendere l'oggetto domanda come primo parametro e la scelta del computer come secondo parametro.

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

Se la scelta del computer corrisponde alla risposta, `getResults` deve restituire `The computer's choice is correct!`

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

Se la scelta del computer non corrisponde alla risposta, `getResults` deve restituire `The computer's choice is wrong. The correct answer is: <correct-answer>`, dove `<correct-answer>` è il valore della risposta corretta alla domanda scelta.

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

La tua funzione `getResults` deve usare il confronto di uguaglianza esatto, non il confronto di sottostringhe.

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
