---
id: 61968f8877c6720d6d61aaf5
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`body` 안에 `div`가 `class`인 `ground`을 추가하세요.

# --hints--

새로운 `div`를 추가해야 합니다.

```js
assert.exists(document.querySelector('div'));
```

`div`에 `class`인 `ground`를 지정해야 합니다.

```js
assert.include(document.querySelector('div')?.className, 'ground');
```

`div`를 `body` 안에 배치해야 합니다.

```js
assert.exists(document.querySelector('body > div.ground'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}
```
