---
id: 619665c9abd72906f3ad30f9
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

행복한 Flappy Penguin을 만들면서 CSS 변형과 애니메이션을 더 연습할 것입니다.

스타일시트를 페이지에 연결하는 것부터 시작하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('link'));
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert.match(code, /<head>[\w\W\s]*<link[\w\W\s]*\/?>[\w\W\s]*<\/head>/i);
```

`link` 요소는 값이 `rel`인 `stylesheet` 속성을 가져야 합니다.

```js
const link_element = document.querySelector('link');
const rel = link_element.getAttribute("rel");
assert.equal(rel, "stylesheet");
```

`link` 요소는 값이 `href`인 `styles.css` 속성을 가져야 합니다.

```js
const link = document.querySelector('link');
assert.equal(link.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
