---
id: 619691693bc14b0e528f5a20
title: Step 7
challengeType: 0
dashedName: step-7
---

# --description--

`.ground` 요소를 선택하고, `width`이 뷰포트 전체 너비를 차지하도록 설정하세요. 그런 다음, `height`를 `400px`로 설정하세요.

# --hints--

`.ground` 선택자를 사용하세요.

```js
assert.match(code, /\.ground\s*\{/);
```

`.ground` 요소에 `width`를 `100vw`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.width, '100vw');
```

`.ground` 요소에 `height`를 `400px`로 지정하세요.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.height, '400px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--
```
