---
id: 61969c487ced6f12db8fef94
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

`.left-mountain` 요소를 선택하고, 그 `width`과 `height`를 `300px`로 설정하세요. 그런 다음, `background`를 `rgb(203, 241, 228)`에서 시작해 `rgb(80, 183, 255)`에서 끝나는 선형 그라디언트로 설정하세요.

# --hints--

`.left-mountain` 선택자를 사용하세요.

```js
assert.match(code, /\.left-mountain\s*\{/);
```

`.left-mountain`에 `width`를 `300px`으로 지정하세요. 예상되는 `--fcc-actual--`은 `--fcc-expected--`입니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.width, '300px');
```

`.left-mountain`에 `height`를 `300px`으로 지정하세요. 예상되는 `--fcc-actual--`은 `--fcc-expected--`입니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left-mountain')?.height, '300px');
```

`.left-mountain`에 `background`를 `linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255))`으로 지정하세요.

```js
assert.include(['linear-gradient(rgb(203,241,228),rgb(80,183,255))', 'rgba(0,0,0,0)linear-gradient(rgb(203,241,228),rgb(80,183,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.left-mountain')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--

--fcc-editable-region--

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
