---
id: 61969aa6acef5b12200f672e
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

배경에 풍경을 만들기 위해 두 개의 산을 추가하세요.

`.penguin` 요소 위에 `div`를 `class`의 `left-mountain`와 함께 추가하세요.

# --hints--

`div` 안에 새로운 `body`를 추가해야 합니다. --fcc-expected--개의 `div` 요소가 있어야 하지만 --fcc-actual--개가 발견되었습니다.

```js
assert.lengthOf(document.querySelectorAll('body > div'), 3);
```

`div`에 `class`인 `left-mountain`를 지정해야 합니다.

```js
assert.include(document.querySelector('body > div:not(.ground, .penguin)')?.className, 'left-mountain');
```

`.left-mountain`를 `.penguin` 앞에 배치해야 합니다.

```js
assert.strictEqual(document.querySelector('.penguin')?.previousElementSibling, document.querySelector('.left-mountain'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

--fcc-editable-region--
  <body>

    <div class="penguin"></div>
    <div class="ground"></div>
  </body>
--fcc-editable-region--
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
}
```
