---
id: 6197cff995d03905b0cca8ad
title: 28단계
challengeType: 0
dashedName: step-28
---

# --description--

펭귄은 머리와 몸 두 개의 주요 섹션으로 구성됩니다.

`.penguin` 안에 두 개의 새로운 `div` 요소를 추가하세요. 첫 번째는 `class`가 `penguin-head`이고, 두 번째는 `class`가 `penguin-body`입니다.

# --hints--

`div`에 두 개의 새로운 `.penguin` 요소를 추가해야 합니다. 예상된 --fcc-expected-- `.penguin > div` 요소가 있지만 실제로는 --fcc-actual--가 발견되었습니다.

```js
assert.lengthOf(document.querySelectorAll('.penguin > div'), 2);
```

첫 번째 `div`에 `class`가 `penguin-head`인 속성을 부여해야 합니다.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(1)')?.className, 'penguin-head');
```

두 번째 `div`에 `class`가 `penguin-body`인 속성을 부여해야 합니다.

```js
assert.include(document.querySelector('.penguin > div:nth-of-type(2)')?.className, 'penguin-body');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="left-mountain"></div>
    <div class="back-mountain"></div>
    <div class="sun"></div>
--fcc-editable-region--
    <div class="penguin">

    </div>
--fcc-editable-region--
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

.left-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(80, 183, 255));
  position: absolute;
  transform: skew(0deg, 44deg);
  z-index: 2;
  margin-top: 100px;
}

.back-mountain {
  width: 300px;
  height: 300px;
  background: linear-gradient(rgb(203, 241, 228), rgb(47, 170, 255));
  position: absolute;
  z-index: 1;
  transform: rotate(45deg);
  left: 110px;
  top: 225px;
}

.sun {
  width: 200px;
  height: 200px;
  background-color: yellow;
  position: absolute;
  border-radius: 50%;
  top: -75px;
  right: -75px;
}

.penguin {
  width: 300px;
  height: 300px;
  margin: auto;
  margin-top: 75px;
}

.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));
  z-index: 3;
  position: absolute;
  margin-top: -58px;
}
```
