---
id: 6197f667297bb30a552ce017
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

`.ground` 요소가 페이지 레이아웃의 쌓임 맥락에서 세 번째가 되도록, `z-index`을 `3`로 설정하고 `position`을 `absolute`로 설정하세요.

# --hints--

`.ground`에 `z-index` 값을 가진 `3`을 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.zIndex, '3');
```

`.ground`에 `position` 값을 가진 `absolute`을 지정해야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.ground')?.position, 'absolute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="./styles.css" />
    <title>Penguin</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>

  <body>
    <div class="ground"></div>
  </body>
</html>
```

```css
body {
  background: linear-gradient(45deg, rgb(118, 201, 255), rgb(247, 255, 222));
  margin: 0;
  padding: 0;
  width: 100%;
  height: 100vh;
  overflow: hidden;
}

--fcc-editable-region--
.ground {
  width: 100vw;
  height: 400px;
  background: linear-gradient(90deg, rgb(88, 175, 236), rgb(182, 255, 255));

}
--fcc-editable-region--
```
