---
id: 68cabc534d2e741f33a3fe1d
title: Criar um Rastreador de Estatísticas de Personagem de Jogo
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

Neste laboratório, você vai criar um rastreador de estatísticas de personagem de jogo. O programa permitirá que você crie um personagem com atributos específicos, atualize esses atributos e recupere as estatísticas atuais do personagem.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Crie uma classe chamada `GameCharacter` que representa um personagem do jogo e gerencia as estatísticas do personagem.
2. Quando instanciado, um novo objeto `GameCharacter` deve ter os seguintes atributos:
    - `_name` definido para a string fornecida no momento da instanciação.
    - `_health` definido como `100`.
    - `_mana` definido como `50`.
    - `_level` definido como `1`.
3. Crie uma propriedade `name` para acesso somente leitura ao nome do personagem.
4. Para a propriedade `health`:
     - Defina um getter que retorne a saúde atual.
     - Defina um setter que impeça que a saúde seja definida abaixo de `0` e limite a saúde em `100`.
5. Para a propriedade `mana`:
     - Defina um getter que retorna a mana atual.
     - Defina um setter que impeça que a mana seja definida abaixo de `0` e limite a mana em `50`.
6. Crie um getter para `level` para retornar o nível atual do personagem.
7. Defina um método chamado `level_up` que:
    - Aumenta o nível do personagem em 1.
    - Redefine a vida para `100` e o mana para `50` usando seus respectivos setters de propriedade.
    - Exibe uma mensagem no formato `<name> leveled up to <level>!` (onde `<name>` e `<level>` devem ser substituídos pelo nome do personagem e pelo novo nível, respectivamente).
8. Defina um método `__str__` que retorne uma string formatada incluindo:
    - O nome do personagem.
    - O nível do personagem.
    - A saúde atual do personagem.
    - O mana atual do personagem.  
Por exemplo, um personagem chamado `Kratos`, logo após a instanciação, deve ser representado como o seguinte:

   ```md
    Name: Kratos
    Level: 1
    Health: 100
    Mana: 50
   ```

## Exemplo de uso

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Você deve ter uma classe `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Você deve ter um método `__init__` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

O método `__init__` deve ter dois parâmetros, sendo o primeiro `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Você deve definir `self._name` para a string fornecida no momento da instanciação dentro do método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Você deve atribuir `100` a `self._health` dentro do método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Você deve atribuir `50` a `self._mana` dentro do método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Você deve atribuir `1` a `self._level` dentro do método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Você deve ter um método `name` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

O método `name` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Você deve retornar `self._name` no método `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

O método `name` deve ter um decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Você deve ter um método `health` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

O método `health` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Você deve retornar `self._health` no método `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

O método `health` deve ter um decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Você deve criar um método setter `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

O método setter `health` deve ter um decorador `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

O método setter `health` deve ter dois parâmetros, sendo o primeiro `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Você deve definir `self._health` como `0` se o valor dado ao setter `health` for menor que `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Você deve definir `self._health` como `100` se o valor dado ao setter `health` for maior que `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Você deve definir `self._health` para o valor dado ao setter `health` quando o valor estiver entre `0` e `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
`)})
```

Você deve ter um método `mana` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

O método `mana` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Você deve retornar `self._mana` no método `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

O método `mana` deve ter um decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Você deve criar um método setter `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

O método setter `mana` deve ter um decorador `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

O método setter `mana` deve ter dois parâmetros, sendo o primeiro `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Você deve definir `self._mana` como `0` se o valor dado ao setter `mana` for menor que `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Você deve definir `self._mana` para `50` se o valor dado ao setter `mana` for maior que `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Você deve definir `self._mana` para o valor dado ao setter `mana` se o valor estiver entre `0` e `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Você deve ter um método `level` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

O método `level` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Você deve retornar `self._level` no método `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

O método `level` deve ter um decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Você deve ter um método `level_up` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

O método `level_up` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Você deve aumentar `self._level` em `1` no método `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Você deve definir `self.health` para `100` no método `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Você deve definir `self.mana` para `50` no método `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

O método `level_up` deve imprimir `<name> leveled up to <level>!` (onde `<name>` e `<level>` devem ser substituídos pelo nome do personagem e pelo novo nível, respectivamente).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Você deve ter um método `__str__` na classe `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

O método `__str__` deve ter apenas um parâmetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Seu método `__str__` deve retornar uma string com as estatísticas do personagem usando o formato fornecido.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
