---
id: 69708d59833ceded7d00541f
title: Implementar Operações de Lista Ligada
challengeType: 26
dashedName: implement-linked-list-operations
---

# --description--

Neste laboratório, você irá implementar operações adicionais para uma estrutura de dados de lista ligada, construindo sobre a funcionalidade básica de adicionar e remover nós.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário**

1. Você deve ter uma função `contains` que aceita uma lista ligada e um elemento. Ela deve retornar `true` se o elemento especificado existir na lista ligada e `false` caso contrário.
1. Você deve ter uma `getAt` função que aceita uma lista ligada e um índice. Ela deve retornar o elemento no índice dado na lista ligada. Se o índice estiver fora dos limites, deve retornar `undefined`.
1. Você deve ter uma função `insertAt` que aceita uma lista ligada, um índice e um elemento. Ela deve inserir o elemento dado na posição especificada na lista ligada. Se o índice estiver fora dos limites, ela não deve modificar a lista.
1. Você deve ter uma `removeAt` função que aceita uma lista ligada e um índice. Ela deve remover o nó no índice fornecido na lista ligada. Se o índice estiver fora dos limites, não deve modificar a lista.
1. Você deve ter uma função `clear` que aceita uma lista ligada. Ela deve remover todos os elementos da lista ligada, efetivamente redefinindo-a para um estado vazio.

**Nota:** Alguns testes posteriores dependem de métodos anteriores. Por exemplo, se `getAt` não for implementado corretamente, os testes para funções como `insertAt` e `removeAt` podem falhar mesmo quando essas funções estiverem quase corretas.

# --hints--

`initList` deve retornar um objeto com `head` definido como `null` e `length` definido como `0`.

```js
const list = initList();
assert.isNull(list.head);
assert.strictEqual(list.length, 0);
```

`add` deve aumentar o comprimento da lista em um.

```js
const list = initList();
add(list, 'apple');
assert.strictEqual(list.length, 1);
```

`remove` deve diminuir o comprimento da lista em um.

```js
const list = initList();  
add(list, 'apple');  
add(list, 'banana');  
remove(list, 'apple');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(list.head.element, 'banana');  
assert.isNull(list.head.next);
```

Você deve ter uma `contains` function.

```js
assert.isFunction(contains);
```

`contains` deve retornar `true` se o elemento existir na lista.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
add(list, 'cherry');
assert.isTrue(contains(list, 'banana'));
```

`contains` deve retornar `false` se o elemento não existir na lista.

```js
const list = initList();
add(list, 'apple');
add(list, 'banana');
assert.isFalse(contains(list, 'cherry'));
```

`contains` deve retornar `false` para uma lista vazia.

```js
const list = initList();
assert.isFalse(contains(list, 'apple'));
```

Você deve ter uma `getAt` função.

```js
assert.isFunction(getAt);
```

`getAt` deve retornar o elemento em um dado índice.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
```

`getAt` deve retornar `undefined` para índices inválidos.

```js
const list = initList();  
add(list, 'apple');  
assert.isUndefined(getAt(list, -1));  
assert.isUndefined(getAt(list, 1));  
assert.isUndefined(getAt(list, 5));
```

`getAt` deve retornar `undefined` para uma lista vazia.

```js
const list = initList();  
assert.isUndefined(getAt(list, 0));
```

Você deve ter uma `insertAt` função.

```js
assert.isFunction(insertAt);
```

`insertAt` deve inserir um elemento no início da lista.

```js
const list = initList();  
add(list, 'second');  
add(list, 'third');  
insertAt(list, 0, 'first');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` deve inserir um elemento no meio da lista.

```js
const list = initList();  
add(list, 'first');  
add(list, 'third');  
insertAt(list, 1, 'second');  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.strictEqual(getAt(list, 2), 'third');  
assert.strictEqual(list.length, 3);
```

`insertAt` deve inserir um elemento no final da lista.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
insertAt(list, 2, 'third');
assert.strictEqual(getAt(list, 0), 'first');
assert.strictEqual(getAt(list, 1), 'second');
assert.strictEqual(getAt(list, 2), 'third');
assert.strictEqual(list.length, 3);
```

`insertAt` deve inserir em uma lista vazia no índice `0`.  

```js  
const list = initList();  
insertAt(list, 0, 'first');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.isNull(list.head.next); 
```

`insertAt` não deve modificar a lista para índices inválidos.

```js
const list = initList();  
add(list, 'apple');  
insertAt(list, -1, 'invalid');  
insertAt(list, 10, 'invalid');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
assert.isUndefined(getAt(list, 1));
```

Você deve ter uma `removeAt` função.

```js
assert.isFunction(removeAt);
```

`removeAt` deve remover o elemento em um índice especificado.

```js
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 1);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'third');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` deve remover o primeiro elemento corretamente.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
removeAt(list, 0);
assert.strictEqual(getAt(list, 0), 'second');
assert.strictEqual(list.length, 1);
assert.isNull(list.head.next);
```

`removeAt` deve remover o último elemento corretamente.  

```js  
const list = initList();  
add(list, 'first');  
add(list, 'second');  
add(list, 'third');  
removeAt(list, 2);  
assert.strictEqual(getAt(list, 0), 'first');  
assert.strictEqual(getAt(list, 1), 'second');  
assert.isUndefined(getAt(list, 2));  
assert.strictEqual(list.length, 2);
```

`removeAt` deve redefinir a lista ao remover o único elemento.

```js
const list = initList();  
add(list, 'only');  
removeAt(list, 0);  
assert.strictEqual(list.length, 0);  
assert.isNull(list.head);  
assert.isTrue(isEmpty(list));
```

`removeAt` não deve modificar a lista para índices inválidos.

```js
const list = initList();  
add(list, 'apple');  
removeAt(list, -1);  
removeAt(list, 10);  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'apple');  
```

Você deve ter uma `clear` função.

```js
assert.isFunction(clear);
```

`clear` deve remover todos os elementos da lista.

```js
const list = initList();
add(list, 'first');
add(list, 'second');
add(list, 'third');
clear(list);
assert.strictEqual(list.length, 0);
assert.isNull(list.head);
assert.isTrue(isEmpty(list));
```

`clear` deve permitir que a lista seja reutilizada.  

```js  
const list = initList();  
add(list, 'first');  
clear(list);  
clear(list);  
add(list, 'second');  
assert.strictEqual(list.length, 1);  
assert.strictEqual(getAt(list, 0), 'second');  
assert.isNull(list.head.next); 
```

# --seed--

## --seed-contents--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  
}

function getAt(list, index) {
 
}

function insertAt(list, index, element) {
  
}

function removeAt(list, index) {
 
}

function clear(list) {

}
```

# --solutions--

```js
function initList() {
  return {
    head: null,
    length: 0
  };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
  const node = { element, next: null };

  if (isEmpty(list)) {
    list.head = node;
  } else {
    let current = list.head;
    while (current.next !== null) {
      current = current.next;
    }
    current.next = node;
  }

  list.length++;
}

function remove(list, element) {
  let previous = null;
  let current = list.head;

  while (current !== null && current.element !== element) {
    previous = current;
    current = current.next;
  }

  if (current === null) return;

  if (previous !== null) {
    previous.next = current.next;
  } else {
    list.head = current.next;
  }

  list.length--;
}

function contains(list, element) {
  let current = list.head;
  while (current !== null) {
    if (current.element === element) {
      return true;
    }
    current = current.next;
  }
  return false;
}

function getAt(list, index) {
  if (index < 0 || index >= list.length) return undefined;
  let current = list.head;
  let i = 0;
  while (i < index) {
    current = current.next;
    i++;
  }
  return current.element;
}

function insertAt(list, index, element) {
  if (index < 0 || index > list.length) return;
  const node = { element, next: null };
  if (index === 0) {
    node.next = list.head;
    list.head = node;
    list.length++;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = node;
  node.next = current;
  list.length++;
}

function removeAt(list, index) {
  if (index < 0 || index >= list.length) return;
  if (index === 0) {
    list.head = list.head.next;
    list.length--;
    return;
  }
  let current = list.head;
  let previous = null;
  let i = 0;
  while (i < index) {
    previous = current;
    current = current.next;
    i++;
  }
  previous.next = current.next;
  list.length--;
}

function clear(list) {
  list.head = null;
  list.length = 0;
}
```
