---
id: 66629f407d679d3105e8317f
title: "Construye un juego de piedra, papel o tijera"
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
Tu juego se jugarﾃ｡ contra la computadora. Escribe una funciﾃｳn que devuelva aleatoriamente `"rock"`, `"paper"` o `"scissors"`.

No necesitas preocuparte por la parte del front-end del juego. Solo escribe la lﾃｳgica del juego. Abre el `script.js` y comienza a seguir las instrucciones.

**Historias de usuario:**

1. Debes tener una funciﾃｳn llamada `getComputerChoice`.

1. Tu funciﾃｳn `getComputerChoice` debe devolver aleatoriamente `"rock"`, `"paper"`, o `"scissors"`.

**Pista:** El mﾃｩtodo [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) devuelve un nﾃｺmero aleatorio que es mayor o igual a 0 y menor que 1. Piensa en cﾃｳmo puedes usar esto para devolver condicionalmente una de las mﾃｺltiples opciones.

Tu juego serﾃ｡ jugado por un jugador humano. Escribe una funciﾃｳn que tome la elecciﾃｳn del usuario y la devuelva.

1. Crea una funciﾃｳn llamada `getHumanChoice`.

1. Escribe el cﾃｳdigo para que `getHumanChoice` devuelva una de las opciones vﾃ｡lidas dependiendo de lo que el usuario ingrese.

**Sugerencia:** Usa el mﾃｩtodo [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) para obtener la entrada del usuario.

Tu juego lleva el registro de la puntuaciﾃｳn del jugador. Escribe variables para llevar el registro de la puntuaciﾃｳn del jugador.

1. Crea dos nuevas variables llamadas `humanScore` y `computerScore` en el ﾃ｡mbito global.

1. Inicializa esas variables con el valor de `0`.

Tu juego se jugarﾃ｡ ronda por ronda. Escribe una funciﾃｳn que tome las elecciones del jugador humano y la computadora como argumentos, juegue una sola ronda, incremente la puntuaciﾃｳn del ganador de la ronda y registre un anuncio del ganador.

1. Crea una nueva funciﾃｳn llamada `playRound`.

1. Define dos parﾃ｡metros para `playRound`. Parﾃ｡metro uno `humanChoice` y parﾃ｡metro dos `computerChoice`. Usa estos dos parﾃ｡metros para tomar las elecciones del humano y del ordenador como argumentos.

1. Haz que el parﾃ｡metro `humanChoice` de tu funciﾃｳn sea insensible a mayﾃｺsculas y minﾃｺsculas para que los jugadores puedan ingresar `"rock"`, `"ROCK"`, `"RocK"`, u otras variaciones.

1. Escribe el cﾃｳdigo para tu funciﾃｳn `playRound` que devuelva un valor de cadena que represente al ganador de la ronda.

- Si hay un empate, debe devolver `"It's a tie!"`.

- Si el jugador gana, deberﾃｭa devolver `"You win! [player choice] beats [computer choice]"`.

- Si la computadora gana, deberﾃｭa devolver `"You lose! [computer choice] beats [player choice]"`.

1. Incrementa la variable `humanScore` o `computerScore` basada en el ganador de la ronda.

Tu juego jugarﾃ｡ 3 rondas. Escribirﾃ｡s una funciﾃｳn llamada `playGame` que llame a `playRound` para jugar 3 rondas, lleve el registro de las puntuaciones y declare un ganador al final.

1. Crea una nueva funciﾃｳn llamada `playGame`.

1. Crea un bucle que juegue 3 rondas y llame a la funciﾃｳn `playRound` cada vez con las funciones de elecciﾃｳn del humano y la computadora como argumentos.

1. Al final del juego, devuelve al ganador del juego basado en quien ganﾃｳ mﾃ｡s rondas.

- Si el jugador humano gana mﾃ｡s rondas que la computadora, devuelve un mensaje que diga `"You win the game!"`.

- Si la computadora gana mﾃ｡s rondas que el jugador humano, devuelve un mensaje que diga `"You lose the game!"`.

# --hints--

Debes tener una funciﾃｳn llamada `getComputerChoice`

```js
assert.isFunction(getComputerChoice);
```

Tu funciﾃｳn `getComputerChoice` debe devolver aleatoriamente `"rock"`, `"paper"`, o `"scissors"`.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

Debes tener una funciﾃｳn llamada `getHumanChoice`

```js
assert.isFunction(getHumanChoice);
```

Debes tener dos variables llamadas `humanScore` y `computerScore` en el ﾃ｡mbito global.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

Debes tener una funciﾃｳn llamada `playRound`

```js
assert.isFunction(playRound);
```

Tu funciﾃｳn `playRound` debe tomar las elecciones del jugador humano y de la computadora como argumentos con los parﾃ｡metros `humanChoice` y `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

Tu funciﾃｳn `playRound` debe ser insensible a mayﾃｺsculas y minﾃｺsculas para que los jugadores puedan ingresar `"rock"`, `"ROCK"`, `"RocK"`, u otras variaciones.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

Tu funciﾃｳn `playRound` debe devolver la cadena `"It's a tie!"` si las elecciones del humano y la computadora son las mismas.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

Tu funciﾃｳn `playRound` deberﾃｭa devolver la cadena `"You win! [player choice] beats [computer choice]"` si el jugador gana.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

Tu funciﾃｳn `playRound` deberﾃｭa devolver la cadena `"You lose! [computer choice] beats [player choice]"` si la computadora gana.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

Tu funciﾃｳn `playRound` debe incrementar la variable `humanScore` o `computerScore` basada en el ganador de la ronda.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

Debes tener una funciﾃｳn llamada `playGame`.

```js
assert.isFunction(playGame);
```

Debes usar un bucle para jugar 3 rondas.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

Debes devolver al ganador del juego basado en quien ganﾃｳ mﾃ｡s rondas.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```

