---
id: 68e6a99a0fd7386fb66e156f
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你将通过创建一个游戏设置窗格来练习为复选框设置样式。

所有的超文本标记语言公式化（`DOCTYPE`、`html`、`head` 和 `body`）都已为你提供。

在 `body` 内，创建一个包含嵌套 `h1` 元素的 `div` 元素。为 `h1` 元素赋予文本 `Game Settings`。

# --hints--

你应该有一个 `div` 元素。

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

你应该有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

你的 `h1` 元素应位于你的 `div` 元素内。

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

你的 `h1` 元素应包含文本 `Game Settings`。

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
