---
id: 68e9a7bd60c971278d2d4ffc
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

在之前的课程中，你学习了 CSS 中的伪类和伪元素。现在你将应用这些知识，通过创建一个带有伪类选择器的组合类型选择器。

首先，通过将其设置为 `input[type="checkbox"]` 并在末尾追加 `:checked`，开始为类型选择器设置形式。

接下来，为它设置一个值为 `#f1be32` 的 `background-color`。

最后，将 `border-color` 设置为 `#e2a60d`。

# --hints--

你应该有一个组合的类型和伪类选择器 `input[type="checkbox"]:checked`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

你应该为你的组合类型和伪类选择器 `input[type="checkbox"]:checked` 添加一个 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

你应该有一个 `background-color` 属性，值为 `#f1be32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

你应该为你的组合类型和伪类选择器 `input[type="checkbox"]:checked` 添加一个 `border-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

你应该有一个 `border-color` 属性，值为 `#e2a60d`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
