---
id: 68e9a3194160f8f5620ddf21
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

紧接着你的 `h1` 元素，创建四个 `label` 元素。

在每个 `label` 元素内部，创建一个 `input` 元素，其 `type` 属性设置为 `"checkbox"`。

在每个 `input` 元素之后，依次为你的 `label` 元素添加以下文本：`Sound Effects`、`Background Music`、`Hard Mode` 和 `Haptic Feedback`。

# --hints--

你应该有四个 `label` 元素。

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

你的第一个 `label` 元素应包含一个 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

你的第一个 `input` 元素的 `type` 属性应设置为 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

你的第一个 `label` 元素应具有文本 `Sound Effects`。

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

你的 `Sound Effects` 文本应位于 `input` 元素之后。

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

你的第二个 `label` 元素应包含一个 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

你的第二个 `input` 元素的 `type` 属性应设置为 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

你的第二个 `label` 元素应包含文本 `Background Music`。

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

你的 `Background Music` 文本应位于 `input` 元素之后。

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

你的第三个 `label` 元素应包含一个 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

你的第三个 `input` 元素的 `type` 属性应设置为 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

你的第三个 `label` 元素应具有 `Hard Mode` 文本。

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

你的 `Hard Mode` 文本应位于 `input` 元素之后。

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

你的第四个 `label` 元素应包含一个 `input` 元素。

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

你的第四个 `input` 元素的 `type` 属性应设置为 `checkbox`。

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

你的第四个 `label` 元素应具有文本 `Haptic Feedback`。

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

你的 `Haptic Feedback` 文本应位于 `input` 元素之后。

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
