---
id: 68e9a6bf1a74721a4e43e217
title: 步骤 10
challengeType: 0
dashedName: step-10
---

# --description--

现在你将移除浏览器应用的默认复选框。在你的 `input[type='checkbox']` 选择器内，添加值为 `none` 的 `appearance`。

将 `appearance` 属性设置为 `none` 将清除浏览器应用于复选框的外观，允许你显示想要的样式。

完成后，由于复选框将不再可见，设置一个带有 `2px` `solid` 和十六进制代码为 `#f1be32` 的 `border`。

# --hints--

你应该为你的 `input[type="checkbox"]` 选择器添加一个 `appearance` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

你应该有一个 `appearance` 属性，值为 `none`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

你应该为你的 `input[type="checkbox"]` 选择器添加一个 `border` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

你应该有一个 `border` 属性，值为 `2px solid #f1be32`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
