---
id: 68e9a834dfef752dc35f6f90
title: 步骤14
challengeType: 0
dashedName: step-14
---

# --description--

现在，当你的复选框在用户点击时改变颜色，你将应用一些样式，使复选框被选中更加明显。

CSS 伪类可以链式连接。你可以在另一个伪类或伪元素后添加一个伪类或伪元素，以定位满足多个条件的元素。

为 `input[type="checkbox"]` 创建一个新选择器。为你的选择器添加 `:checked` 伪类，后跟 `::after` 伪元素。这样你就可以在复选框被选中时添加一个视觉指示器。

接下来，为你的 `content` 添加值为 `"✓"`。

# --hints--

你应该有一个组合类型选择器，包含伪类选择器和伪元素选择器 `input[type="checkbox"]:checked::after`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

你应该为你的组合类型选择器添加一个带有伪类选择器和伪元素选择器 `input[type="checkbox"]:checked::after` 的 `content` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

你应该有一个 `content` 属性，值为 `"✓"`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```

