---
id: 68e9a54dd97fd9091e18acce
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

现在创建一个 `h1` 选择器并将 `text-align` 属性设置为 `center`。这将使你的 `h1` 元素居中。

# --hints--

你的 CSS 中应该有一个用于 `h1` 元素的类型选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("h1"));
```

你应该为你的 `h1` 选择器添加一个 `text-align` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign);
```

你应该有一个 `text-align` 属性，值为 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("h1")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

--fcc-editable-region--

--fcc-editable-region--

```
