---
id: 68ee86e3ee07b01626a392cc
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在你的 `.settings-card` 选择器内将 `margin` 属性设置为 `auto`。

将 `margin` 属性设置为 `auto` 会自动调整元素的边距，以均匀分配其容器中剩余的空间，通常用于水平居中块级元素。

最后，设置一个值为 `left` 的 `text-align` 属性。这将使内联内容，例如文本，左对齐到其包含元素的左侧。

# --hints--

你应该为你的 `settings-card` 类选择器添加一个 `margin` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

你应该有一个 `margin` 属性，值为 `auto`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

你应该为 `settings-card` 类选择器添加一个 `text-align` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

你应该有一个 `text-align` 属性，值为 `left`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
