---
id: 68f15cd397503429f0345262
title: 步骤15
challengeType: 0
dashedName: step-15
---

# --description--

在 `input[type="checkbox"]:checked::after` 选择器的声明中，将 `display` 设置为 `block`。将你的 `display` 属性设置为 `block` 会使元素成为块级元素，这意味着它占据其容器的全部宽度并从新行开始，从而使你更容易控制其宽度、高度和间距。

接下来，将 `text-align` 设置为 `center`。这将使内联内容（如文本或内联元素）在块内水平居中。块仍将占据其容器的全部宽度，但其中的所有内容都会对齐到中心。

# --hints--

你应该为你的组合类型选择器添加一个 `display` 属性，该选择器带有伪类选择器和伪元素选择器 `input[type="checkbox"]:checked::after`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

你应该有一个 `display` 属性，值为 `block`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

你应该为你的组合类型选择器添加一个 `text-align` 属性，该选择器带有伪类选择器和伪元素选择器 `input[type="checkbox"]:checked::after`。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

你应该有一个 `text-align` 属性，值为 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
