---
id: 68f15ef0446464325ce1aa45
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

你需要在复选框和标签之间留出一些间距。首先在你的 CSS 中使用 `label` 选择器，并设置一个 `display` 属性为 `block`。

接下来，设置一个 `margin` 属性为 `8px auto`，以添加元素之间的垂直间距。这将为每个元素的上下创建一致的间距，有助于分隔内容并提高可读性。

最后，将 `cursor` 属性设置为 `pointer` 于 `label` 元素上。当用户悬停在 `label` 上时，这将把光标变成手形图标，表示该元素可点击，从而提升整体用户体验。

# --hints--

你的 CSS 中应该有一个用于 `label` 的类型选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

你应该为你的 `label` 选择器添加一个 `display` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

你应该有一个 `display` 属性，值为 `block`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

你应该为你的 `label` 选择器添加 `margin` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

你应该有一个 `margin` 属性，值为 `8px auto`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

你应该为你的 `label` 选择器添加一个 `cursor` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

你应该有一个值为 `pointer` 的 `cursor` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
