---
id: 68f1560aff1214103ab6b59b
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

现在你将开始用一些 CSS 样式美化页面。首先创建 `body` 选择器。

将 `body` 设置为具有值为 `100vh` 的 `height` 属性和值为 `#f0f0f0` 的 `background-color` 属性。

`100vh` 的 `height` 使 `body` 占据浏览器视口的全部高度，而浅灰色背景为页面提供了一个细腻的基础。

最后，设置一个值为 `center` 的 `text-align` 属性。这样会使页面内所有内联内容居中，除非子元素用自己的对齐方式覆写它。

# --hints--

你应该使用 `body` 类型选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

你应该在 `body` 类型选择器中添加 `height` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

你应该有一个 `height` 属性，值为 `100vh`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

你应该在 `body` 类型选择器内添加 `background-color` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

你应该有一个 `background-color` 属性，值为 `#f0f0f0`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

你应该在 `body` 类型选择器内添加 `text-align` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

你应该有一个 `text-align` 属性，值为 `center`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
