---
id: 68f29dbf9125476b4808d6fb
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

现在是为 `settings-card` 容器设置样式的时候了。为 `settings-card` 创建一个类选择器。

这里将是你所有容器格式的所在。将 `max-width` 设置为 `250px` 以定义你的容器的整体大小。

接下来，将你的 `padding` 设置为 `20px`，以便你的内容与容器的边框之间有间距。

之后，通过将你的 `border-radius` 设置为 `10px` 来创建圆角。

然后设置一个 `box-shadow`，其值为 `0 2px 6px rgba(0,0,0,0.2)`。这将创建一个细微的“悬浮”效果，为容器营造深度感。

# --hints--

你的 css 中应该有一个用于 `settings-card` 的类选择器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

你应该为 `.settings-card` 类选择器添加 `max-width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

你应该有一个 `max-width` 属性，值为 `250px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

你应该为 `.settings-card` 类选择器添加一个 `padding` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

你应该有一个 `padding` 属性，值为 `20px`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

你应该为你的 `.settings-card` 类选择器添加 `border-radius` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

你应该有一个值为 `10px` 的 `border-radius` 属性。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

你应该为你的 `.settings-card` 类选择器添加一个 `box-shadow` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

你应该有一个值为 `0 2px 6px rgba(0,0,0,0.2)` 的 `box-shadow`。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
