---
id: 66c64095f61166ede6e39a84
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Para esta lista de compras, utilizará un array para representar los artículos que necesita comprar.

En lecciones anteriores, aprendiste cómo crear arreglos así:

```js
const fruits = ['apple', 'banana', 'orange'];
```

En este paso, crea una variable llamada `shoppingList` y asígnale un array vacío.

# --hints--

Debe tener una variable llamada `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Debe asignar un array a la variable `shoppingList`.

```js
assert.isArray(shoppingList);
```

El arreglo debería estar vacío.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
