---
id: 66c64210413532ee9d3bd342
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Para esta siguiente parte del taller, practicará agregando artículos alimenticios a la lista de compras.

Comience utilizando `console.log` para registrar el mensaje `"It will be nice to have some fruit to eat."`. 

# --hints--

Debería registrar el mensaje `"It will be nice to have some fruit to eat."` en la consola.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
