---
id: 66c645b345a39bf04864dc50
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Para ver la lista de compras actualizada, deberá registrar la lista de compras actual y un mensaje corto en la consola.

Dado que este mensaje se utilizará repetidamente durante el taller, es mejor crear una función reutilizable.

Cree una función llamada `getShoppingListMsg` que tome un array como parámetro y devuelva la cadena `"Lista de Compras Actual: "` seguida por el contenido del array proporcionado.

Puedes usar literales de plantilla o concatenación de cadenas con el operador `+` para combinar la cadena con el arreglo pasado a la función.

# --hints--

Debería tener una función llamada `getShoppingListMsg` en su código.

```js
assert.isFunction(getShoppingListMsg);
```

Su función `getShoppingListMsg` debería devolver una cadena.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Su función `getShoppingListMsg` debería aceptar un array como argumento.


```js
assert.strictEqual(getShoppingListMsg(["Anchovies", "Olives"]), "Current Shopping List: Anchovies,Olives");
```

Su función `getShoppingListMsg` debería aceptar un array y devolver el mensaje `"Lista de Compras Actual: "` seguido por el contenido de ese array. Verifique dos veces los espacios y los errores ortográficos.

```js
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
