---
id: 66c72b0ffbc5522525768558
title: Paso 11
challengeType: 1
dashedName: step-11
---

# --description--

En pasos anteriores, revisaste cómo agregar un elemento al final del array utilizando el método `push`.

Pero el método `push` acepta múltiples argumentos, por lo que puedes agregar varios elementos al final del array de esta manera:

```js
array.push(item1, item2, item3);
```

En este paso, usa el método `push` para añadir las cadenas `"Popcorn", "Beef Jerky", "Potato Chips"` al array `shoppingList`.

El orden es importante, así que asegúrate de añadir los elementos en el orden en que están listados.

# --hints--

Tu array `shoppingList` debe tener un total de `6` cadenas.

```js
assert.lengthOf(shoppingList, 6);
```

Tu `shoppingList` debe tener la cadena `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Tu `shoppingList` debe tener la cadena `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Tu `shoppingList` debe tener la cadena `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
