---
id: 67172b43f84bcd2dec238a3d
title: 책 정리 도구 만들기
challengeType: 26
dashedName: build-a-book-organizer
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `books`라는 이름의 객체 배열이 있어야 하며, 배열의 각 객체는 문자열 `title`, 또 다른 문자열 `authorName`, 그리고 숫자 `releaseYear`를 가져야 합니다.

2. `books` 배열에는 최소 세 개의 객체가 있어야 합니다.

3. 배열을 정렬하기 위해 두 권의 책을 매개변수로 받는 `sortByYear`라는 콜백 함수가 있어야 합니다.

4. `sortByYear` 함수는 첫 번째 책의 `releaseYear`가 두 번째 책의 `releaseYear`보다 작으면 `-1`를 반환해야 합니다.

5. `sortByYear` 함수는 첫 번째 책의 `releaseYear`가 두 번째 책의 `releaseYear`보다 크면 `1`를 반환해야 합니다.

6. `sortByYear` 함수는 두 `releaseYear` 값이 같으면 `0`를 반환해야 합니다.

7. `books` 배열에서 1950년 이후에 출판된 책을 필터링하여 새로운 배열 `filteredBooks`에 저장해야 합니다.

8. `filteredBooks` 배열의 책들을 `releaseYear`에 따라 오름차순으로 정렬해야 합니다. 이전 수업에서 `sort()` 메서드가 배열을 제자리에서 정렬한다는 것을 배웠습니다. 즉, `filteredBooks` 배열이 변경됩니다.

# --hints--

코드에 `sortByYear` 함수가 있어야 합니다.

```js
assert.isFunction(sortByYear);
```

`sortByYear` 함수는 두 개의 매개변수를 받아야 합니다.

```js
assert.lengthOf(sortByYear, 2);
```

`sortByYear` 함수는 `book1` 객체의 `releaseYear`가 `book2` 객체의 `releaseYear`보다 작으면 `-1`를, `book1` 객체의 `releaseYear`가 `book2` 객체의 `releaseYear`보다 크면 `1`를, 그 외 모든 경우에는 `0`를 반환해야 합니다.

```js
assert.equal(sortByYear({releaseYear: 1913}, {releaseYear: 1925}), -1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1913}), 1);
assert.equal(sortByYear({releaseYear: 1925}, {releaseYear: 1925}), 0);
```

코드에 `books` 배열이 있어야 합니다.

```js
assert.isArray(books)
```

`books` 배열에는 최소 세 개의 책 객체가 있어야 합니다.

```js
assert.isAtLeast(books.length, 3)
```

`books` 배열에는 문자열 `title`, 또 다른 문자열 `authorName`, 그리고 숫자 `releaseYear`를 가진 객체가 포함되어야 합니다.

```js

books.forEach((book) => {
  assert.isObject(book);
  assert.sameMembers(Object.keys(book), ["title", "authorName", "releaseYear"]);
  assert.isString(book.title);
  assert.isNotEmpty(book.title);
  assert.isString(book.authorName);
  assert.isNotEmpty(book.authorName);
  assert.isNumber(book.releaseYear);
})
```

코드에 `filteredBooks` 배열이 있어야 합니다.

```js
assert.isArray(filteredBooks)
```

`filteredBooks` 배열은 `books` 배열에서 일부 책이 필터링된 것이어야 하며 비어 있으면 안 됩니다.

```js
assert.isBelow(filteredBooks.length, books.length)
assert.isNotEmpty(filteredBooks)
```

`filteredBooks` 배열에는 원래 `books` 배열에 있던 책만 포함되어야 합니다.

```js
filteredBooks.forEach(book => {
  assert.isTrue(
    books.some(original => 
      original.title === book.title &&
      original.authorName === book.authorName &&
      original.releaseYear === book.releaseYear
    )
  );
});
```

`filteredBooks` 배열에는 지정된 연도 이전 또는 같은 해에 출판된 책만 포함되어야 합니다.

```js
const maxYear = Math.max(...filteredBooks.map(book => book.releaseYear));
const removedYears = books.filter(book => !filteredBooks.includes(book)).map(book => book.releaseYear);
assert.isTrue(removedYears.every(year => year > maxYear));
```

`sort` 고차 함수에 `sortByYear` 콜백 함수를 `filteredBooks` 배열에 전달하여 호출해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /\s*filteredBooks\s*\.\s*sort\s*\(\s*sortByYear\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sortByYear(book1, book2) {
  if (book1.releaseYear < book2.releaseYear) {
    return -1;
  }

  if (book1.releaseYear > book2.releaseYear) {
    return 1;
  }

  return 0;
}

const books = [
  {
    title: "The Great Gatsby",
    authorName: "F. Scott Fitzgerald",
    releaseYear: 1925,
  },
  {
    title: "In Search of Lost Time",
    authorName: "Marcel Proust",
    releaseYear: 1913,
  },
  {
    title: "Ulysses",
    authorName: "James Joyce",
    releaseYear: 1922,
  },
  {
    title: "One Hundred Years of Solitude",
    authorName: "Gabriel García Márquez",
    releaseYear: 1967,
  },
  {
    title: "The Catcher in the Rye",
    authorName: "J. D. Salinger",
    releaseYear: 1951,
  },
];

let filteredBooks = books.filter((book) => book.releaseYear < 1950);

filteredBooks.sort(sortByYear);

console.log("Books Written Before 1950 (sorted according to release year)");
filteredBooks.forEach((book) => {
  console.log(`${book.title} by ${book.authorName} (${book.releaseYear})`);
});
```
