---
id: 5dc174fcf86c76b9248c6eb2
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將通過創建一個貓咪照片應用程序繼續使用基本的 HTML 元素，如標題、段落和列表。

通過添加一個文本爲 `CatPhotoApp` 的 `h1` 元素，開始工作坊。

# --hints--

代碼中應包含文本 `CatPhotoApp`。 你可能需要檢查你的拼寫。

```js
assert.match(code, /catphotoapp/i);
```

你的 `h1` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標籤。 結束標籤的語法爲：`</elementName>`。

```js
assert.match(code, /<\/h1\>/);
```

你的 `h1` 元素的文本應該是 `CatPhotoApp`。 你要麼遺漏了文本或拼寫有誤，要麼沒有把文本放在 `h1` 元素的開始和結束標籤之間。

```js
assert.equal(document.querySelector('h1')?.innerText.toLowerCase(), 'catphotoapp');
```


# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
