---
id: 5dc17d3bf86c76b9248c6eb4
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `h2` 元素下方創建一個 `p` 元素，併爲其添加以下文本：

`Everyone loves cute cats online!`

# --hints--

你的 `p` 元素應該有一個開始標籤。 開始標籤的語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('p'));
```

Your `p` tags should be in lowercase. By convention, all HTML tags are written in lowercase.

```js
assert.notMatch(code, /<\/?P>/);
```

你的 `p` 元素應該有一個結束標籤。 結束標籤在 `<` 字符之後有一個 `/`。

```js
assert.match(code, /<\/p\>/);
```

你的 `p` 元素的文本應爲 `Everyone loves cute cats online!`。你要麼遺漏了文本，要麼輸入有誤。

```js
const extraSpacesRemoved = document
  .querySelector('p')
  ?.innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /everyone loves cute cats online!$/i);
```

你的 `p` 元素應該在 `h2` 元素下面。 你把它們的順序寫錯了。

```js
const collection = [...document.querySelectorAll('h2,p')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('H2'), collection.indexOf('P'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
--fcc-editable-region--
    <h2>Cat Photos</h2>

--fcc-editable-region--
  </body>
</html>
```
