---
id: 587d7db8367417b2b2512ba0
title: مطابقة كل شيء ما عدا الحروف والأرقام
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

تعلمت أنه يمكنك استخدام اختصار لمطابقة الأحرف الأبجدية الرقمية `[A-Za-z0-9_]` باستخدام `\w`. نمط طبيعي قد ترغب في البحث عنه هو عكس الأحرف الأبجدية الرقمية.

يمكنك البحث عن عكس `\w` باستخدام `\W`. لاحظ أن النمط العكسي يستخدم حرفًا كبيرًا. هذا الاختصار يعادل `[^A-Za-z0-9_]`.

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

ستُرجع أول استدعاء `match` القيمة `["%"]` والثاني سيُرجع `["!"]`.

# --instructions--

استخدم فئة الحروف المختصرة `\W` لحساب عدد الأحرف غير الأبجدية الرقمية في اقتباسات وسلاسل نصية مختلفة.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك العلامة العامة.

```js
assert(nonAlphabetRegex.global);
```

يجب أن يجد التعبير النمطي الخاص بك 6 أحرف غير أبجدية رقمية في السلسلة `The five boxing wizards jump quickly.`.

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

يجب أن يستخدم التعبير النمطي الخاص بك الحرف المختصر لمطابقة الأحرف غير الأبجدية الرقمية.

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

يجب أن يجد التعبير النمطي الخاص بك 8 أحرف غير أبجدية رقمية في السلسلة `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

يجب أن يجد التعبير النمطي الخاص بك 6 أحرف غير أبجدية رقمية في السلسلة `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

يجب أن يجد التعبير النمطي الخاص بك 12 حرفًا غير أبجدي رقمي في السلسلة `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
