---
id: 587d7db7367417b2b2512b9f
title: مطابقة كل الحروف والأرقام
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

باستخدام أصناف الأحرف، استطعت البحث عن كل حروف الأبجدية باستخدام `[a-z]`. هذا النوع من أصناف الأحرف شائع بما يكفي ليكون له اختصار، رغم أنه يشمل بعض الأحرف الإضافية أيضًا.

أقرب صنف أحرف في JavaScript لمطابقة الأبجدية هو `\w`. هذا الاختصار يعادل `[A-Za-z0-9_]`. هذا الصنف يطابق الحروف الكبيرة والصغيرة بالإضافة إلى الأرقام. لاحظ أن هذا الصنف يشمل أيضًا حرف الشرطة السفلية (`_`).

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

كل هذه الاستدعاءات الأربعة لـ `test` سترجع `true`.

تُعرف هذه الأصناف المختصرة أيضًا باسم <dfn>أصناف الأحرف المختصرة</dfn>.

# --instructions--

استخدم صنف الأحرف المختصر `\w` لحساب عدد الأحرف الأبجدية الرقمية في اقتباسات وسلاسل نصية مختلفة.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك العلامة العامة (global flag).

```js
assert(alphabetRegexV2.global);
```

يجب أن يستخدم التعبير النمطي الخاص بك حرف الاختصار `\w` لمطابقة كل الأحرف التي هي أبجدية رقمية.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

يجب أن يجد التعبير النمطي 31 حرفًا أبجديًا رقميًا في السلسلة النصية `The five boxing wizards jump quickly.`

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

يجب أن يجد التعبير النمطي 32 حرفًا أبجديًا رقميًا في السلسلة النصية `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

يجب أن يجد التعبير النمطي 30 حرفًا أبجديًا رقميًا في السلسلة النصية `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

يجب أن يجد التعبير النمطي 36 حرفًا أبجديًا رقميًا في السلسلة النصية `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
