---
id: 587d7db5367417b2b2512b94
title: مطابقة أي شيء باستخدام النقطة العامة
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

أحيانًا لن تعرف (أو لا تحتاج إلى معرفة) الأحرف الدقيقة في أنماطك. التفكير في كل الكلمات التي تطابق، مثلًا، خطأ إملائي سيستغرق وقتًا طويلًا. لحسن الحظ، يمكنك توفير الوقت باستخدام الحرف العام: `.`

الحرف العام `.` يطابق أي حرف واحد. يُسمى الحرف العام أيضًا `dot` و `period`. يمكنك استخدام الحرف العام مثل أي حرف آخر في التعبير النمطي. على سبيل المثال، إذا أردت مطابقة `hug` و `huh` و `hut` و `hum`، يمكنك استخدام التعبير النمطي `/hu./` لمطابقة الكلمات الأربع كلها.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

كلا الاستدعاءين `test` سيُرجعان `true`.

# --instructions--

أكمل التعبير النمطي `unRegex` بحيث يطابق السلاسل `run` و `sun` و `fun` و `pun` و `nun` و `bun`. يجب أن يستخدم التعبير النمطي الحرف العام.

# --hints--

يجب أن تستخدم الطريقة `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

يجب أن تستخدم الحرف العام في التعبير النمطي `unRegex`

```js
assert(/\./.test(unRegex.source));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `run` في السلسلة `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `sun` في السلسلة `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `fun` في السلسلة `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `pun` في السلسلة `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `nun` في السلسلة `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

يجب أن يطابق التعبير النمطي `unRegex` كلمة `bun` في السلسلة `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

يجب ألا يطابق التعبير النمطي `unRegex` السلسلة `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

يجب ألا يطابق التعبير النمطي `unRegex` السلسلة `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
