---
id: 587d7db4367417b2b2512b92
title: استخراج التطابقات
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

حتى الآن، كنت تتحقق فقط مما إذا كان نمط موجودًا أم لا داخل سلسلة نصية. يمكنك أيضًا استخراج التطابقات الفعلية التي وجدتها باستخدام الطريقة `.match()`.

لاستخدام الطريقة `.match()`، طبق الطريقة على سلسلة نصية ومرر التعبير النمطي داخل الأقواس العادية.

إليك مثالًا:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

هنا ستُرجع أول `match` القيمة `["Hello"]` والثاني سيُرجع `["expressions"]`.

لاحظ أن صيغة `.match` هي "عكس" الطريقة `.test` التي كنت تستخدمها حتى الآن:

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

طبق الطريقة `.match()` لاستخراج السلسلة النصية `coding`.

# --hints--

يجب أن يحتوي `result` على السلسلة النصية `coding`

```js
assert(result.join() === 'coding');
```

يجب أن يبحث التعبير النمطي `codingRegex` عن السلسلة النصية `coding`

```js
assert(codingRegex.source === 'coding');
```

يجب أن تستخدم الطريقة `.match()`.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
