---
id: 587d7db8367417b2b2512ba3
title: مطابقة الفراغات
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

حتى الآن، غطت التحديات مطابقة حروف الأبجدية والأرقام. يمكنك أيضًا مطابقة الفراغات أو المسافات بين الحروف.

يمكنك البحث عن الفراغات باستخدام `\s`، وهي حرف `s` صغير. هذا النمط لا يطابق الفراغات فقط، بل أيضًا عودة العربة، الجدولة، تغذية النموذج، وحروف السطر الجديد. يمكنك اعتباره مشابهًا لصنف الحروف `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

سيُرجع هذا الاستدعاء `match` القيمة `[" ", " "]`.
# --instructions--

غيّر التعبير النمطي `countWhiteSpace` ليبحث عن عدة حروف فراغ في سلسلة نصية.

# --hints--

يجب أن يستخدم التعبير النمطي العلامة العامة (global flag).

```js
assert(countWhiteSpace.global);
```

يجب أن يستخدم التعبير النمطي الحرف المختصر `\s` لمطابقة كل حروف الفراغ.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

يجب أن يجد التعبير النمطي ثمانية فراغات في السلسلة النصية `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

يجب أن يجد التعبير النمطي ثلاثة فراغات في السلسلة النصية `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

يجب أن لا يجد التعبير النمطي أي فراغات في السلسلة النصية `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
