---
id: 587d7db8367417b2b2512ba1
title: مطابقة كل الأحرف غير الرقمية
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

أظهر التحدٍّ الأخير كيفية البحث عن الأرقام باستخدام الاختصار `\d` مع حرف `d` صغير. يمكنك أيضًا البحث عن الأحرف غير الرقمية باستخدام اختصار مشابه يستخدم حرف `D` كبير بدلاً من ذلك.

الاختصار للبحث عن الأحرف غير الرقمية هو `\D`. هذا يعادل صنف الأحرف `[^0-9]`، الذي يبحث عن حرف واحد ليس رقمًا بين الصفر والتسعة.

# --instructions--

استخدم صنف الأحرف المختصر للأحرف غير الرقمية `\D` لحساب عدد الأحرف غير الرقمية في عناوين الأفلام.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك حرف الاختصار لمطابقة الأحرف غير الرقمية

```js
assert(/\\D/.test(noNumRegex.source));
```

يجب أن يستخدم التعبير النمطي الخاص بك العلامة العامة.

```js
assert(noNumRegex.global);
```

يجب أن لا يجد التعبير النمطي الخاص بك أي أحرف غير رقمية في السلسلة `9`.

```js
assert('9'.match(noNumRegex) == null);
```

يجب أن يجد التعبير النمطي الخاص بك 6 أحرف غير رقمية في السلسلة `Catch 22`.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

يجب أن يجد التعبير النمطي الخاص بك 11 حرفًا غير رقمي في السلسلة `101 Dalmatians`.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

يجب أن يجد التعبير النمطي الخاص بك 15 حرفًا غير رقمي في السلسلة `One, Two, Three`.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

يجب أن يجد التعبير النمطي الخاص بك 12 حرفًا غير رقمي في السلسلة `21 Jump Street`.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

يجب أن يجد التعبير النمطي الخاص بك 17 حرفًا غير رقمي في السلسلة `2001: A Space Odyssey`.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
