---
id: 587d7db6367417b2b2512b9b
title: إيجاد الأحرف باستخدام المطابقة الكسولة
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

في التعبيرات النمطية، المطابقة <dfn>الجشعة</dfn> تجد أطول جزء ممكن من سلسلة نصية يتناسب مع نمط التعبير النمطي وتُرجعه كمطابقة. البديل يسمى المطابقة <dfn>الكسولة</dfn>، التي تجد أصغر جزء ممكن من السلسلة النصية يفي بنمط التعبير النمطي.

يمكنك تطبيق التعبير النمطي `/t[a-z]*i/` على السلسلة النصية `"titanic"`. هذا التعبير النمطي هو في الأساس نمط يبدأ بـ `t`، وينتهي بـ `i`، ويحتوي على بعض الحروف بينهما.

التعبيرات النمطية تكون جشعة افتراضيًا، لذا ستُرجع المطابقة `["titani"]`. فهي تجد أكبر جزء فرعي ممكن يتناسب مع النمط.

مع ذلك، يمكنك استخدام الحرف `?` لتغييره إلى المطابقة الكسولة. عند مطابقة `"titanic"` مع التعبير النمطي المعدل `/t[a-z]*?i/`، تُرجع المطابقة `["ti"]`.

**ملاحظة:** يجب تجنب تحليل HTML باستخدام التعبيرات النمطية، لكن مطابقة نمط سلسلة HTML باستخدام التعبيرات النمطية أمر مقبول تمامًا.

# --instructions--

صحح التعبير النمطي `/<.*>/` ليُرجع وسم HTML `<h1>` وليس النص `"<h1>Winter is coming</h1>"`. تذكر أن الرمز البدل `.` في التعبير النمطي يطابق أي حرف.

# --hints--

يجب أن تكون قيمة المتغير `result` مصفوفة تحتوي على `<h1>`

```js
assert(result[0] == '<h1>');
```

يجب أن يستخدم `myRegex` المطابقة الكسولة

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

يجب ألا يتضمن `myRegex` السلسلة النصية `h1`

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
