---
id: 587d7db4367417b2b2512b93
title: العثور على أكثر من التطابق الأول
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

حتى الآن، كنت قادرًا فقط على استخراج أو البحث عن نمط مرة واحدة.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

هنا سترجع `match` القيمة `["Repeat"]`.

للبحث أو استخراج نمط أكثر من مرة، يمكنك استخدام علم البحث العام: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

وهنا ترجع `match` القيمة `["Repeat", "Repeat", "Repeat"]`

# --instructions--

باستخدام التعبير النمطي `starRegex`، ابحث واستخرج كلمتي `Twinkle` من السلسلة النصية `twinkleStar`.

**ملاحظة**  
يمكنك استخدام عدة أعلام في التعبير النمطي مثل `/search/gi`

# --hints--

يجب أن يستخدم التعبير النمطي `starRegex` العلم العام `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

يجب أن يستخدم التعبير النمطي `starRegex` العلم الذي يتجاهل حالة الأحرف `i`

```js
assert(starRegex.flags.match(/i/).length == 1);
```

يجب أن يطابق البحث كلا الحالتين لكلمة `Twinkle`

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

يجب أن يحتوي ناتج البحث `result` على عنصرين.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
