---
id: 587d7dba367417b2b2512ba8
title: تحقق من الكل أو لا شيء
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

أحيانًا تحتوي الأنماط التي تريد البحث عنها على أجزاء قد توجد أو لا توجد. مع ذلك، قد يكون من المهم التحقق منها رغم ذلك.

يمكنك تحديد وجود عنصر محتمل باستخدام علامة الاستفهام، `?`. هذا يتحقق من وجود صفر أو واحد من العنصر السابق. يمكنك اعتبار هذا الرمز كأنه يشير إلى أن العنصر السابق اختياري.

على سبيل المثال، هناك اختلافات بسيطة بين الإنجليزية الأمريكية والإنجليزية البريطانية، ويمكنك استخدام علامة الاستفهام لمطابقة كلا التهجيتين.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

كلا الاستخدامين لطريقة `test` سيُرجعان `true`.

# --instructions--

غيّر التعبير النمطي `favRegex` ليطابق كلًا من النسخة الأمريكية (`favorite`) والنسخة البريطانية (`favourite`) من الكلمة.

# --hints--

يجب أن يستخدم التعبير النمطي رمز الاختيار، `?`.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
