---
id: 587d7db4367417b2b2512b90
title: مطابقة سلسلة نصية حرفية مع احتمالات مختلفة
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

باستخدام التعبيرات النمطية مثل `/coding/`، يمكنك البحث عن النمط `coding` في سلسلة نصية أخرى.

هذا مفيد للبحث في سلاسل نصية واحدة، لكنه محدود بنمط واحد فقط. يمكنك البحث عن أنماط متعددة باستخدام عامل `alternation` أو عامل `OR`: `|`.

هذا العامل يطابق الأنماط التي تسبقه أو تلحق به. على سبيل المثال، إذا أردت مطابقة السلاسل النصية `yes` أو `no`، فإن التعبير النمطي الذي تحتاجه هو `/yes|no/`.

يمكنك أيضًا البحث عن أكثر من نمطين. تفعل ذلك بإضافة المزيد من الأنماط مفصولة بعوامل `OR` إضافية، مثل `/yes|no|maybe/`.

# --instructions--

أكمل التعبير النمطي `petRegex` ليطابق الحيوانات الأليفة `dog`، `cat`، `bird`، أو `fish`.

# --hints--

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `true` للسلسلة النصية `John has a pet dog.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `false` للسلسلة النصية `Emma has a pet rock.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `true` للسلسلة النصية `Emma has a pet bird.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `true` للسلسلة النصية `Liz has a pet cat.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `false` للسلسلة النصية `Kara has a pet dolphin.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `true` للسلسلة النصية `Alice has a pet fish.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

يجب أن يُرجع التعبير النمطي `petRegex` القيمة `false` للسلسلة النصية `Jimmy has a pet computer.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
