---
id: 587d7db7367417b2b2512b9e
title: مطابقة أنماط نهاية السلاسل النصية
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

في التحدٍّ السابق، تعلمت استخدام حرف مؤشر الإدراج للبحث عن أنماط في بداية السلاسل النصية. هناك أيضًا طريقة للبحث عن أنماط في نهاية السلاسل النصية.

يمكنك البحث في نهاية السلاسل النصية باستخدام رمز علامة الدولار `$` في نهاية التعبير النمطي.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

ستُعيد أول استدعاء للدالة `test` القيمة `true`، بينما سيُعيد الثاني القيمة `false`.

# --instructions--

استخدم وسم الرابط أو وسم الرابط للنص العالي (`$`) لمطابقة السلسلة النصية `caboose` في نهاية السلسلة النصية `caboose`.

# --hints--

يجب أن تبحث عن `caboose` باستخدام وسم الرابط أو وسم الرابط للنص العالي `$` في التعبير النمطي الخاص بك.

```js
assert(lastRegex.source == 'caboose$');
```

يجب ألا يستخدم التعبير النمطي أي علامات.

```js
assert(lastRegex.flags == '');
```

يجب أن تطابق `caboose` في نهاية السلسلة النصية `The last car on a train is the caboose`.

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
