---
id: 587d7db3367417b2b2512b8e
title: استخدام طريقة Test
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

تُستخدم التعبيرات النمطية في لغات البرمجة لمطابقة أجزاء من السلاسل النصية. تُنشئ أنماطًا لمساعدتك في إجراء تلك المطابقة.

إذا أردت العثور على الكلمة `the` في السلسلة النصية `The dog chased the cat`، يمكنك استخدام التعبير النمطي التالي: `/the/`. لاحظ أن علامات الاقتباس غير مطلوبة داخل التعبير النمطي.

تحتوي JavaScript على عدة طرق لاستخدام التعبيرات النمطية. إحدى الطرق لاختبار تعبير نمطي هي استخدام الطريقة `.test()`. تأخذ الطريقة `.test()` التعبير النمطي، وتطبقه على سلسلة نصية (توضع داخل الأقواس العادية)، وتُرجع `true` أو `false` إذا وجد نمطك شيئًا أم لا.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

تعيد الطريقة `test` هنا القيمة `true`.

# --instructions--

طبق التعبير النمطي `myRegex` على السلسلة النصية `myString` باستخدام الطريقة `.test()`.

# --hints--

يجب عليك استخدام `.test()` لاختبار التعبير النمطي.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

يجب أن تُرجع النتيجة الخاصة بك القيمة `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
