---
id: 587d7db8367417b2b2512ba2
title: تقييد أسماء المستخدمين المحتملة
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

تُستخدم أسماء المستخدمين في كل مكان على الإنترنت. هي التي تمنح المستخدمين هوية فريدة على مواقعهم المفضلة.

تحتاج إلى التحقق من جميع أسماء المستخدمين في قاعدة بيانات. إليك بعض القواعد البسيطة التي يجب على المستخدمين اتباعها عند إنشاء اسم المستخدم.

1) يمكن لأسماء المستخدمين استخدام الأحرف الأبجدية الرقمية فقط.

2) يجب أن تكون الأرقام الوحيدة في اسم المستخدم في النهاية. يمكن أن يكون هناك صفر أو أكثر منها في النهاية. لا يمكن أن يبدأ اسم المستخدم برقم.

3) يمكن أن تكون حروف اسم المستخدم أحرفًا صغيرة أو كبيرة.

4) يجب أن يكون طول أسماء المستخدمين حرفين على الأقل. يمكن لاسم المستخدم المكون من حرفين أن يستخدم فقط حروف الأبجدية كأحرف.

# --instructions--

غيّر التعبير النمطي `userCheck` ليتناسب مع القيود المذكورة أعلاه.

# --hints--

يجب أن يطابق التعبير النمطي السلسلة النصية `JACK`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `J`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Jo`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Oceans11`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `RegexGuru`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `007`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `9`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `A1`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `BadUs3rnam3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Z97`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `c57bT3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `AB1`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `J%4`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
