---
id: 587d7db3367417b2b2512b8f
title: مطابقة السلاسل النصية الحرفية
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

في التحدٍّ الأخير، بحثت عن الكلمة `Hello` باستخدام التعبير النمطي `/Hello/`. هذا التعبير النمطي بحث عن تطابق حرفي للسلسلة النصية `Hello`. إليك مثالًا آخر يبحث عن تطابق حرفي للسلسلة النصية `Kevin`:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

سيُرجع هذا الاستدعاء `test` القيمة `true`.

أي أشكال أخرى من `Kevin` لن تتطابق. على سبيل المثال، التعبير النمطي `/Kevin/` لن يتطابق مع `kevin` أو `KEVIN`.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

سيُرجع هذا الاستدعاء `test` القيمة `false`.

سيُظهر تحدٍّ مستقبلي كيفية مطابقة تلك الأشكال الأخرى أيضًا.

# --instructions--

أكمل التعبير النمطي `waldoRegex` لإيجاد `"Waldo"` في السلسلة النصية `waldoIsHiding` بتطابق حرفي.

# --hints--

يجب أن يجد التعبير النمطي `waldoRegex` السلسلة النصية `Waldo`

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

يجب ألا يبحث التعبير النمطي `waldoRegex` عن أي شيء آخر.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

يجب أن تُجري تطابقًا حرفيًا للسلسلة النصية باستخدام التعبير النمطي.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
