---
id: 587d7db7367417b2b2512b9d
title: مطابقة أنماط بداية السلاسل النصية
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

أظهرت التحديات السابقة أن التعبيرات النمطية تُستخدم للبحث عن عدد من التطابقات. كما تُستخدم أيضًا للبحث عن أنماط في مواقع محددة داخل السلاسل النصية.

في تحدٍّ سابق، استخدمت رمز مؤشر الإدراج (`^`) داخل مجموعة أحرف لإنشاء مجموعة أحرف منفية بالشكل `[^thingsThatWillNotBeMatched]`. خارج مجموعة الأحرف، يُستخدم مؤشر الإدراج للبحث عن أنماط في بداية السلاسل النصية.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

ستُعيد أول استدعاء للدالة `test` القيمة `true`، بينما سيُعيد الثاني القيمة `false`.

# --instructions--

استخدم رمز مؤشر الإدراج في تعبير نمطي للعثور على `Cal` فقط في بداية السلسلة النصية `rickyAndCal`.

# --hints--

يجب أن يبحث التعبير النمطي عن السلسلة النصية `Cal` بحرف كبير.

```js
assert(calRegex.source == '^Cal');
```

يجب ألا يستخدم التعبير النمطي أي علامات.

```js
assert(calRegex.flags == '');
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Cal` في بداية السلسلة.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Cal` في وسط السلسلة.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
