---
id: 587d7db6367417b2b2512b98
title: مطابقة الأحرف المفردة غير المحددة
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

حتى الآن، أنشأت مجموعة من الأحرف التي تريد مطابقتها، لكن يمكنك أيضًا إنشاء مجموعة من الأحرف التي لا تريد مطابقتها. تُسمى هذه الأنواع من مجموعات الأحرف <dfn>مجموعات الأحرف المنفية</dfn>.

لإنشاء مجموعة أحرف منفية، ضع رمز مؤشر الإدراج (`^`) بعد القوس المفتوح وقبل الأحرف التي لا تريد مطابقتها.

على سبيل المثال، `/[^aeiou]/gi` تطابق كل الأحرف التي ليست حرفًا متحركًا. لاحظ أن الأحرف مثل `.`, `!`, `[`, `@`, `/` والمسافات البيضاء تُطابق - مجموعة الأحرف المنفية للحروف المتحركة تستثني فقط حروف العلة.

# --instructions--

أنشئ تعبيرًا نمطيًا واحدًا يطابق كل الأحرف التي ليست رقمًا أو حرفًا متحركًا. تذكر تضمين العلامات المناسبة في التعبير النمطي.

# --hints--

يجب أن يطابق التعبير النمطي `myRegex` تسعة عناصر.

```js
assert(result.length == 9);
```

يجب أن يستخدم التعبير النمطي `myRegex` العلامة العامة (global flag).

```js
assert(myRegex.flags.match(/g/).length == 1);
```

يجب أن يستخدم التعبير النمطي `myRegex` العلامة التي تتجاهل حالة الأحرف (case insensitive flag).

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
