---
id: 587d7dbb367417b2b2512baa
title: إعادة استخدام الأنماط باستخدام مجموعات الالتقاط
challengeType: 1
forumTopicId: 301364
dashedName: reuse-patterns-using-capture-groups
---

# --description--

افترض أنك تريد مطابقة كلمة تتكرر عدة مرات كما في المثال أدناه.

```js
let repeatStr = "row row row your boat";
```

يمكنك استخدام `/row row row/`، لكن ماذا لو لم تكن تعرف الكلمة المحددة المتكررة؟ يمكن استخدام <dfn>مجموعات الالتقاط</dfn> للعثور على سلاسل فرعية مكررة.

تُنشأ مجموعات الالتقاط عن طريق إحاطة نمط التعبير النمطي المراد التقاطه بأقواس عادية. في هذه الحالة، الهدف هو التقاط كلمة تتكون من أحرف أبجدية رقمية، لذا ستكون مجموعة الالتقاط هي `\w+` محاطة بأقواس عادية: `/(\w+)/`.

يُحفظ الجزء المطابق بواسطة المجموعة في "متغير" مؤقت، يمكن الوصول إليه داخل نفس التعبير النمطي باستخدام الشرطة المائلة العكسية ورقم مجموعة الالتقاط (مثلاً `\1`). تُرقم مجموعات الالتقاط تلقائيًا حسب موقع أقواسها الافتتاحية (من اليسار إلى اليمين)، بدءًا من 1.

المثال أدناه يطابق كلمة تتكرر ثلاث مرات مفصولة بمسافات:

```js
let repeatRegex = /(\w+) \1 \1/;
repeatRegex.test(repeatStr); // Returns true
repeatStr.match(repeatRegex); // Returns ["row row row", "row"]
```

استخدام الطريقة `.match()` على سلسلة نصية يعيد مصفوفة تحتوي على الجزء المطابق، مع مجموعات الالتقاط الخاصة به.


# --instructions--

استخدم مجموعات الالتقاط في `reRegex` لمطابقة سلسلة تتكون فقط من نفس الرقم مكررًا ثلاث مرات بالضبط مفصولًا بمسافات مفردة.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك فئة الأحرف المختصرة للأرقام.

```js
assert(reRegex.source.match(/\\d/));
```

يجب أن يعيد التعبير النمطي الخاص بك استخدام مجموعة التقاط مرتين.

```js
assert(reRegex.source.match(/\\1|\\2/g).length >= 2);
```

يجب أن يطابق التعبير النمطي الخاص بك السلسلة `42 42 42`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('42 42 42'));
```

يجب أن يطابق التعبير النمطي الخاص بك السلسلة `100 100 100`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('100 100 100'));
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `42 42 42 42`.

```js
assert.equal('42 42 42 42'.match(reRegex.source), null);
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `42 42`.

```js
assert.equal('42 42'.match(reRegex.source), null);
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `101 102 103`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('101 102 103'));
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `1 2 3`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('1 2 3'));
```

يجب أن يطابق التعبير النمطي الخاص بك السلسلة `10 10 10`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('10 10 10'));
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `42\t42\t42`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42\t42\t42'));
```

يجب ألا يطابق التعبير النمطي الخاص بك السلسلة `42  42  42`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42  42  42'));
```

# --seed--

## --seed-contents--

```js
let repeatNum = "42 42 42";
let reRegex = /change/; // Change this line
let result = reRegex.test(repeatNum);
```

# --solutions--

```js
let repeatNum = "42 42 42";
let reRegex = /^(\d+) \1 \1$/;
let result = reRegex.test(repeatNum);
```
