---
id: 587d7dbb367417b2b2512bac
title: إزالة الفراغات من البداية والنهاية
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

أحيانًا تكون هناك حروف فراغ حول السلاسل النصية غير مرغوب فيها لكنها موجودة. المعالجة النموذجية للسلاسل النصية هي إزالة الفراغات من بدايتها ونهايتها.

# --instructions--

اكتب تعبيرًا نمطيًا واستخدم الطرق المناسبة للسلاسل النصية لإزالة الفراغات من بداية ونهاية السلاسل.

**ملاحظة:** طريقة `String.prototype.trim()` ستعمل هنا، لكن عليك إكمال هذا التحدٍّ باستخدام التعبيرات النمطية.

# --hints--

يجب أن تكون قيمة `result` مساوية للسلسلة النصية `Hello, World!`

```js
assert(result === 'Hello, World!');
```

يجب ألا تستخدم في حلك طريقة `String.prototype.trim()`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

يجب ألا يتم تعيين المتغير `result` مباشرة إلى سلسلة نصية

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

يجب ألا تتغير قيمة المتغير `hello`.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
