---
id: 587d7db5367417b2b2512b96
title: مطابقة حروف الأبجدية
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

رأيت كيف يمكنك استخدام <dfn>مجموعات الأحرف</dfn> لتحديد مجموعة من الأحرف للمطابقة، لكن هذا يتطلب كتابة كثيرة عندما تحتاج إلى مطابقة نطاق واسع من الأحرف (مثل كل حرف في الأبجدية). لحسن الحظ، هناك ميزة مدمجة تجعل هذا الأمر قصيرًا وبسيطًا.

داخل مجموعة الأحرف، يمكنك تحديد نطاق من الأحرف للمطابقة باستخدام رمز الشرطة `-`.

على سبيل المثال، لمطابقة الأحرف الصغيرة من `a` إلى `e` تستخدم `[a-e]`.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

بالترتيب، ستُرجع ثلاث استدعاءات `match` القيم `["cat"]`، `["bat"]`، و`null`.

# --instructions--

طابق كل الحروف في السلسلة النصية `quoteSample`.

**ملاحظة**: تأكد من مطابقة الأحرف الكبيرة والصغيرة معًا.

# --hints--

يجب أن يطابق التعبير النمطي `alphabetRegex` 35 عنصرًا.

```js
assert(result.length == 35);
```

يجب أن يستخدم التعبير النمطي `alphabetRegex` العلامة العالمية.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

يجب أن يستخدم التعبير النمطي `alphabetRegex` العلامة التي تتجاهل حالة الأحرف.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
