---
id: 587d7db9367417b2b2512ba5
title: تحديد العدد الأعلى والأدنى للتطابقات
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

تذكّر أنك تستخدم علامة الجمع `+` للبحث عن حرف واحد أو أكثر، والنجمة `*` للبحث عن صفر أو أكثر من الأحرف. هذه الطرق مريحة، لكن أحيانًا تريد مطابقة نطاق معين من الأنماط.

يمكنك تحديد العدد الأدنى والأعلى للأنماط باستخدام <dfn>محددات الكمية</dfn>. تُستخدم محددات الكمية مع الأقواس المعقوفة (`{` و `}`). تضع رقمين بين الأقواس المعقوفة - للعدد الأدنى والأعلى للأنماط.

على سبيل المثال، لمطابقة الحرف `a` فقط عندما يظهر بين `3` و `5` مرات في السلسلة النصية `ah`، سيكون التعبير النمطي الخاص بك `/a{3,5}h/`.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

ستُعيد أول استدعاء `test` القيمة `true`، بينما سيُعيد الثاني القيمة `false`.

# --instructions--

غيّر التعبير النمطي `ohRegex` ليطابق العبارة الكاملة `Oh no` فقط عندما تحتوي على من `3` إلى `6` أحرف `h`.

# --hints--

يجب أن يستخدم التعبير النمطي الأقواس المعقوفة.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Ohh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Ohhh no`

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Ohhhh no`

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Ohhhhh no`

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Ohhhhhh no`

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Ohhhhhhh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
