---
id: 587d7db9367417b2b2512ba6
title: تحديد العدد الأدنى فقط من التطابقات
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

يمكنك تحديد العدد الأدنى والأقصى للنماذج باستخدام محددات الكمية مع الأقواس المعقوفة. أحيانًا تريد فقط تحديد العدد الأدنى للنماذج بدون حد أعلى.

لتحديد العدد الأدنى فقط للنماذج، احتفظ بالرقم الأول متبوعًا بفاصلة.

على سبيل المثال، لمطابقة السلسلة النصية `hah` حيث يظهر الحرف `a` على الأقل `3` مرات، سيكون التعبير النمطي الخاص بك `/ha{3,}h/`.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

بالترتيب، ستُرجع ثلاث استدعاءات `test` القيم `true`، `false`، و `true`.

# --instructions--

غيّر التعبير النمطي `haRegex` ليطابق الكلمة `Hazzah` فقط عندما تحتوي على أربعة أحرف `z` أو أكثر.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك الأقواس المعقوفة.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Hazzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Hazzzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Hazzzzah`

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Hazzzzzah`

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Hazzzzzzah`

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Hazzah` التي تحتوي على 30 حرف `z`.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
