---
id: 587d7dbb367417b2b2512bab
title: استخدام مجموعات الالتقاط للبحث والاستبدال
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

البحث مفيد. ومع ذلك، يمكنك جعل البحث أكثر قوة عندما يغير (أو يستبدل) النص الذي تطابقه أيضًا.

يمكنك البحث واستبدال النص في سلسلة نصية باستخدام `.replace()` على السلسلة. المدخلات لـ `.replace()` هي أولاً نمط التعبير النمطي الذي تريد البحث عنه. المعامل الثاني هو السلسلة التي تستبدل التطابق بها أو دالة تقوم بشيء ما.

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

ستُعيد استدعاء `replace` السلسلة النصية `The sky is blue.`.

يمكنك أيضًا الوصول إلى مجموعات الالتقاط في سلسلة الاستبدال باستخدام علامات الدولار (`$`).

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

ستُعيد استدعاء `replace` السلسلة النصية `Camp Code`.

# --instructions--

اكتب تعبيرًا نمطيًا `fixRegex` يستخدم ثلاث مجموعات التقاط للبحث عن كل كلمة في السلسلة النصية `one two three`. ثم حدّث متغير `replaceText` لاستبدال `one two three` بالسلسلة النصية `three two one` وعيّن النتيجة إلى المتغير `result`. تأكد من استخدام مجموعات الالتقاط في سلسلة الاستبدال باستخدام صيغة علامة الدولار (`$`).

# --hints--

يجب أن تستخدم `.replace()` للبحث والاستبدال.

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

يجب أن يغير التعبير النمطي السلسلة النصية `one two three` إلى السلسلة النصية `three two one`

```js
assert(result === 'three two one');
```

يجب ألا تغيّر السطر الأخير.

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

يجب أن يستخدم `fixRegex` ثلاث مجموعات التقاط على الأقل.

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

يجب أن يستخدم `replaceText` سلاسل التطابق الفرعي المحاطة بأقواس (أي أن سلسلة التطابق الفرعي رقم n، $n، تتوافق مع مجموعة الالتقاط رقم n).

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
