---
id: 587d7dba367417b2b2512ba9
title: النظر للأمام الإيجابي والسلبي
challengeType: 1
forumTopicId: 301360
dashedName: positive-and-negative-lookahead
---

# --description--

<dfn>النظر للأمام</dfn> هي أنماط تخبر JavaScript بالنظر للأمام في سلسلتك النصية للتحقق من أنماط أبعد في النص. هذا مفيد عندما تريد البحث عن أنماط متعددة في نفس السلسلة النصية.

هناك نوعان من النظر للأمام: <dfn>النظر للأمام الإيجابي</dfn> و<dfn>النظر للأمام السلبي</dfn>.

النظر للأمام الإيجابي يتحقق من وجود العنصر في نمط البحث، لكنه لا يطابقه فعليًا. يُستخدم النظر للأمام الإيجابي بالشكل `(?=...)` حيث `...` هو الجزء المطلوب الذي لا يُطابق.

أما النظر للأمام السلبي فيتحقق من عدم وجود العنصر في نمط البحث. يُستخدم النظر للأمام السلبي بالشكل `(?!...)` حيث `...` هو النمط الذي لا تريد وجوده. يُعاد بقية النمط إذا لم يكن جزء النظر للأمام السلبي موجودًا.

النظر للأمام قد يكون محيرًا قليلًا لكن بعض الأمثلة ستساعدك.

```js
let quit = "qu";
let noquit = "qt";
let quRegex= /q(?=u)/;
let qRegex = /q(?!u)/;
quit.match(quRegex);
noquit.match(qRegex);
```

كلا هذين الاستدعاءين لـ `match` سيُعيدان `["q"]`.

استخدام أكثر عملية للنظر للأمام هو التحقق من نمطين أو أكثر في سلسلة نصية واحدة. هنا مدقق كلمة مرور بسيط (بشكل بدائي) يبحث عن كلمة بين 3 و6 أحرف وعلى الأقل رقم واحد:

```js
let password = "abc123";
let checkPass = /(?=\w{3,6})(?=\D*\d)/;
checkPass.test(password);
```

# --instructions--

استخدم النظر للأمام في `pwRegex` لمطابقة كلمات المرور التي تزيد عن 5 أحرف وتحتوي على رقمين متتاليين.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك اثنين من `lookaheads` الإيجابية.

```js
assert(pwRegex.source.match(/\(\?=.*?\)\(\?=.*?\)/) !== null);
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `astronaut`

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('astronaut'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `banan1`

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('banan1'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `bana12`

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('bana12'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `abc123`

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('abc123'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `12345`

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('12345'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `8pass99`

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('8pass99'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `1a2bcde`

```js
pwRegex.lastIndex = 0;
assert(!pwRegex.test('1a2bcde'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `astr1on11aut`

```js
pwRegex.lastIndex = 0;
assert(pwRegex.test('astr1on11aut'));
```

# --seed--

## --seed-contents--

```js
let sampleWord = "astronaut";
let pwRegex = /change/; // Change this line
let result = pwRegex.test(sampleWord);
```

# --solutions--

```js
let pwRegex =  /(?=\w{6})(?=\w*\d{2})/;
```
