---
id: 587d7db5367417b2b2512b97
title: مطابقة الأرقام وحروف الأبجدية
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

استخدام الشرطة (`-`) لمطابقة نطاق من الأحرف لا يقتصر على الحروف فقط. بل يعمل أيضًا لمطابقة نطاق من الأرقام.

على سبيل المثال، `/[0-5]/` تطابق أي رقم بين `0` و `5`، بما في ذلك `0` و `5`.

أيضًا، يمكن دمج نطاق من الحروف والأرقام في مجموعة أحرف واحدة.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

أنشئ تعبيرًا نمطيًا واحدًا يطابق نطاقًا من الحروف بين `h` و `s`، ونطاقًا من الأرقام بين `2` و `6`. تذكر تضمين العلامات المناسبة في التعبير النمطي.

# --hints--

يجب أن يطابق التعبير النمطي `myRegex` 17 عنصرًا.

```js
assert(result.length == 17);
```

يجب أن يستخدم التعبير النمطي `myRegex` العلامة العالمية.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

يجب أن يستخدم التعبير النمطي `myRegex` علامة تجاهل حالة الأحرف.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
