---
id: 587d7db5367417b2b2512b95
title: مطابقة حرف واحد مع احتمالات متعددة
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

تعلمت كيف تطابق الأنماط الحرفية (`/literal/`) والحرف العام (`/./`). هذان هما أقصى حدود التعبيرات النمطية، حيث تجد في أحدهما التطابقات الدقيقة وفي الآخر تطابق كل شيء. هناك خيارات تمثل توازنًا بين هذين الحدين.

يمكنك البحث عن نمط حرفي مع بعض المرونة باستخدام <dfn>فئات الأحرف</dfn>. تسمح فئات الأحرف بتعريف مجموعة من الأحرف التي تريد مطابقتها بوضعها داخل أقواس مربعة (`[` و`]`).

على سبيل المثال، تريد مطابقة `bag` و`big` و`bug` ولكن ليس `bog`. يمكنك إنشاء التعبير النمطي `/b[aiu]g/` للقيام بذلك. الجزء `[aiu]` هو فئة الأحرف التي ستطابق فقط الأحرف `a` أو `i` أو `u`.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

بالترتيب، ستُرجع استدعاءات `match` الأربعة القيم `["big"]` و`["bag"]` و`["bug"]` و`null`.

# --instructions--

استخدم فئة أحرف تحتوي على الحروف المتحركة (`a` و`e` و`i` و`o` و`u`) في التعبير النمطي `vowelRegex` للعثور على كل الحروف المتحركة في السلسلة النصية `quoteSample`.

**ملاحظة:** تأكد من مطابقة الحروف المتحركة الكبيرة والصغيرة.

# --hints--

يجب أن تجد كل الحروف المتحركة الـ 25.

```js
assert(result.length == 25);
```

يجب أن يستخدم التعبير النمطي `vowelRegex` فئة أحرف.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

يجب أن يستخدم التعبير النمطي `vowelRegex` العلامة العامة (global flag).

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

يجب أن يستخدم التعبير النمطي `vowelRegex` العلامة التي تتجاهل حالة الأحرف (case insensitive flag).

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

يجب ألا يطابق التعبير النمطي أي من الحروف الساكنة.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
