---
id: 587d7db7367417b2b2512b9c
title: العثور على مجرم واحد أو أكثر في مطاردة
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

حان الوقت لتتوقف وتختبر مهاراتك الجديدة في كتابة التعبيرات النمطية. هربت مجموعة من المجرمين من السجن وفرت، لكنك لا تعرف عددهم. مع ذلك، تعرف أنهم يبقون قريبين من بعضهم عندما يكونون حول أشخاص آخرين. أنت مسؤول عن العثور على جميع المجرمين دفعة واحدة.

إليك مثالًا لمراجعة كيفية القيام بذلك:

التعبير النمطي `/z+/` يطابق الحرف `z` عندما يظهر مرة واحدة أو أكثر على التوالي. سيجد تطابقات في جميع السلاسل النصية التالية:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

لكنه لا يجد تطابقات في السلاسل النصية التالية لأنه لا توجد حروف `z`:

```js
""
"ABC"
"abcabc"
```

# --instructions--

اكتب تعبيرًا نمطيًا جشعًا يجد مجرمًا واحدًا أو أكثر ضمن مجموعة من الأشخاص الآخرين. يُمثل المجرم بالحرف الكبير `C`.

# --hints--

يجب أن يطابق التعبير النمطي مجرمًا واحدًا (`C`) في السلسلة النصية `C`

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

يجب أن يطابق التعبير النمطي مجرمين اثنين (`CC`) في السلسلة النصية `CC`

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

يجب أن يطابق التعبير النمطي ثلاثة مجرمين (`CCC`) في السلسلة النصية `P1P5P4CCCcP2P6P3`.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

يجب أن يطابق التعبير النمطي خمسة مجرمين (`CCCCC`) في السلسلة النصية `P6P2P7P4P5CCCCCP3P1`

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

يجب ألا يطابق التعبير النمطي أي مجرمين في السلسلة النصية الفارغة `""`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

يجب ألا يطابق التعبير النمطي أي مجرمين في السلسلة النصية `P1P2P3`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

يجب أن يطابق التعبير النمطي خمسين مجرمًا (`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`) في السلسلة النصية `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3`.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
