---
id: 587d7db4367417b2b2512b91
title: تجاهل حالة الأحرف أثناء المطابقة
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

حتى الآن، نظرت إلى التعبيرات النمطية (regexes) لمطابقة السلاسل النصية حرفيًا. لكن أحيانًا، قد تريد أيضًا مطابقة اختلافات حالة الأحرف.

الحالة (أو أحيانًا حالة الحروف) هي الفرق بين الأحرف الكبيرة والأحرف الصغيرة. أمثلة على الأحرف الكبيرة هي `A`، `B`، و `C`. وأمثلة على الأحرف الصغيرة هي `a`، `b`، و `c`.

يمكنك مطابقة الحالتين باستخدام ما يسمى بعلم (flag). هناك أعلام أخرى لكن هنا ستركز على العلم الذي يتجاهل الحالة - العلم `i`. يمكنك استخدامه بإضافته إلى التعبير النمطي. مثال على استخدام هذا العلم هو `/ignorecase/i`. هذا التعبير النمطي يمكنه مطابقة السلاسل النصية `ignorecase`، `igNoreCase`، و `IgnoreCase`.

# --instructions--

اكتب تعبيرًا نمطيًا `fccRegex` لمطابقة `freeCodeCamp` بغض النظر عن حالة الأحرف. يجب ألا يطابق التعبير النمطي أي اختصارات أو تغييرات تحتوي على مسافات.

# --hints--

يجب أن يطابق التعبير النمطي السلسلة النصية `freeCodeCamp`

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FreeCodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FreecodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FreeCodecamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Free Code Camp`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FreeCOdeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `FCC`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FrEeCoDeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FrEeCodECamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `FReeCodeCAmp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
