---
id: 587d7db9367417b2b2512ba7
title: تحديد العدد الدقيق للتطابقات
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

يمكنك تحديد الحد الأدنى والأقصى لعدد الأنماط باستخدام محددات الكمية داخل الأقواس المعقوفة. أحيانًا تريد عددًا محددًا فقط من التطابقات.

لتحديد عدد معين من الأنماط، ضع ذلك الرقم فقط بين الأقواس المعقوفة.

على سبيل المثال، لمطابقة الكلمة `hah` فقط مع الحرف `a` مكرراً `3` مرات، سيكون التعبير النمطي الخاص بك هو `/ha{3}h/`.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

بالتسلسل، ستُرجع ثلاث استدعاءات `test` القيم `false`، `true`، و`false`.

# --instructions--

غيّر التعبير النمطي `timRegex` ليطابق الكلمة `Timber` فقط عندما تحتوي على أربعة أحرف `m`.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك الأقواس المعقوفة.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Timber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Timmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Timmmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

يجب أن يطابق التعبير النمطي السلسلة النصية `Timmmmber`

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

يجب ألا يطابق التعبير النمطي السلسلة النصية `Timber` التي تحتوي على 30 حرفًا `m`.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
