---
id: 587d7db6367417b2b2512b9a
title: مطابقة الأحرف التي تظهر صفر أو أكثر من المرات
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

استخدم التحدٍّ الأخير علامة الجمع `+` للبحث عن الأحرف التي تظهر مرة واحدة أو أكثر. هناك خيار آخر يطابق الأحرف التي تظهر صفر أو أكثر من المرات.

الحرف المستخدم لذلك هو النجمة أو العلامة النجمية: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

بالترتيب، ستُرجع استدعاءات `match` الثلاثة القيم `["goooooooo"]`، `["g"]`، و `null`.

# --instructions--

أنشئ تعبيرًا نمطيًا `chewieRegex` يستخدم الحرف `*` لمطابقة حرف `A` كبير يتبعه مباشرة صفر أو أكثر من أحرف `a` صغيرة في `chewieQuote`. لا يحتاج تعبيرك النمطي إلى أعلام أو فئات أحرف، ويجب ألا يطابق أيًا من الاقتباسات الأخرى.

# --hints--

يجب أن يستخدم تعبيرك النمطي `chewieRegex` الحرف `*` لمطابقة صفر أو أكثر من أحرف `a`.

```js
assert(/\*/.test(chewieRegex.source));
```

يجب أن يطابق تعبيرك النمطي السلسلة النصية `A` في `chewieQuote`.

```js
assert(result[0][0] === 'A');
```

يجب أن يطابق تعبيرك النمطي السلسلة النصية `Aaaaaaaaaaaaaaaa` في `chewieQuote`.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

يجب أن يطابق تعبيرك النمطي `chewieRegex` 16 حرفًا في `chewieQuote`.

```js
assert(result[0].length === 16);
```

يجب ألا يطابق تعبيرك النمطي أي أحرف في السلسلة النصية `He made a fair move. Screaming about it can't help you.`

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

يجب ألا يطابق تعبيرك النمطي أي أحرف في السلسلة النصية `Let him have it. It's not wise to upset a Wookiee.`

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
