---
id: 587d7db6367417b2b2512b99
title: مطابقة الأحرف التي تظهر مرة واحدة أو أكثر
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

أحيانًا، تحتاج إلى مطابقة حرف (أو مجموعة أحرف) يظهر مرة واحدة أو أكثر على التوالي. هذا يعني أنه يظهر مرة واحدة على الأقل، وقد يتكرر.

يمكنك استخدام الحرف `+` للتحقق من ذلك. تذكر، يجب أن يكون الحرف أو النمط موجودًا بشكل متتالي. أي يجب أن يتكرر الحرف واحدًا تلو الآخر.

على سبيل المثال، `/a+/g` سيجد تطابقًا واحدًا في `abc` ويُرجع `["a"]`. وبسبب `+`، سيجد أيضًا تطابقًا واحدًا في `aabc` ويُرجع `["aa"]`.

إذا كان يتحقق من السلسلة `abab`، فسيجد تطابقين ويُرجع `["a", "a"]` لأن أحرف `a` ليست متتالية - هناك حرف `b` بينهما. وأخيرًا، بما أنه لا يوجد `a` في السلسلة `bcd`، فلن يجد تطابقًا.

# --instructions--

تريد إيجاد التطابقات عندما يظهر الحرف `s` مرة واحدة أو أكثر في `Mississippi`. اكتب تعبيرًا نمطيًا يستخدم علامة `+`.

# --hints--

يجب أن يستخدم التعبير النمطي `myRegex` علامة `+` لمطابقة حرف `s` مرة واحدة أو أكثر.

```js
assert(/\+/.test(myRegex.source));
```

يجب أن يطابق التعبير النمطي `myRegex` عنصرين.

```js
assert(result.length == 2);
```

يجب أن تكون المتغير `result` مصفوفة تحتوي على تطابقين من `ss`

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
