---
id: 587d7db9367417b2b2512ba4
title: مطابقة الأحرف غير البيضاء
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

تعلمت البحث عن الفراغات باستخدام `\s`، مع `s` أحرف صغيرة. يمكنك أيضًا البحث عن كل شيء ما عدا الفراغات.

ابحث عن الأحرف غير البيضاء باستخدام `\S`، وهي `s` أحرف كبيرة. هذا النمط لن يطابق الفراغات، أو عودة المؤشر، أو الجدولة، أو تغذية النموذج، أو أحرف السطر الجديد. يمكنك اعتباره مشابهًا لصنف الأحرف `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

القيمة التي تُرجعها الطريقة `.length` ستكون `32`.

# --instructions--

غيّر التعبير النمطي `countNonWhiteSpace` ليبحث عن عدة أحرف غير بيضاء في سلسلة نصية.

# --hints--

يجب أن يستخدم التعبير النمطي الخاص بك العلامة العامة.

```js
assert(countNonWhiteSpace.global);
```

يجب أن يستخدم التعبير النمطي الخاص بك الحرف المختصر `\S` لمطابقة كل الأحرف غير البيضاء.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

يجب أن يجد التعبير النمطي الخاص بك 35 حرفًا غير فراغ في السلسلة `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

يجب أن يجد التعبير النمطي الخاص بك 23 حرفًا غير فراغ في السلسلة `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

يجب أن يجد التعبير النمطي الخاص بك 21 حرفًا غير فراغ في السلسلة `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
