---
id: 5c3dda8b4d8df89bea71600f
title: التحقق من التجميع المختلط للأحرف
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

أحيانًا نريد التحقق من مجموعات الأحرف باستخدام تعبير نمطي، ولتحقيق ذلك نستخدم الأقواس العادية `()`.

إذا أردت إيجاد إما `Penguin` أو `Pumpkin` في سلسلة نصية، يمكنك استخدام التعبير النمطي التالي: `/P(engu|umpk)in/g`

ثم تحقق مما إذا كانت مجموعات السلاسل النصية المطلوبة موجودة في السلسلة الاختبارية باستخدام الطريقة `test()`.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

الطريقة `test` هنا سترجع `true`.

# --instructions--

صحح التعبير النمطي بحيث يتحقق من أسماء `Franklin Roosevelt` أو `Eleanor Roosevelt` مع مراعاة حساسية حالة الأحرف، ويجب أن يسمح بالأسماء الوسطى.

ثم صحح الشفرة بحيث يتم اختبار التعبير النمطي الذي أنشأته مقابل `myString` ويرجع إما `true` أو `false` حسب ما إذا كان التعبير النمطي يطابق.

# --hints--

يجب أن يرجع التعبير النمطي `myRegex` القيمة `true` للسلسلة النصية `Franklin D. Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

يجب أن يرجع التعبير النمطي `myRegex` القيمة `true` للسلسلة النصية `Eleanor Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

يجب أن يرجع التعبير النمطي `myRegex` القيمة `false` للسلسلة النصية `Franklin Rosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

يجب أن يرجع التعبير النمطي `myRegex` القيمة `false` للسلسلة النصية `Frank Roosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

يجب أن يرجع التعبير النمطي `myRegex` القيمة `false` للسلسلة النصية `FranklinRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

يجب أن يرجع التعبير النمطي `myRegex` القيمة `false` للسلسلة النصية `EleanorRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

يجب أن تستخدم `.test()` لاختبار التعبير النمطي.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

يجب أن ترجع النتيجة `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
