---
id: 6ac0b35f3b29a27f7e207295
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

在前一步，你已經創建了變數 `sentence`。現在在醒目標示的那一行輸入 `print(sentence)`，將它的值列印到終端機。

這樣，你就完成了你的第一個 Python 程式。

# --hints--

你的程式碼中應該有 `print(sentence)`。確保從該行的開頭開始輸入。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
