---
id: 6ac09f8420b4b67e075d2cd2
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將創建你的第一個 Python 程式。

這是你的程式碼編輯器，你會在這裡撰寫 Python 程式碼。第 1 行已被凸顯標示。前往該行並輸入：

```py
'Hello'
```

確保從行首開始輸入，並且不要忘記引號。

完成後，點擊「檢查你的程式碼」按鈕，看看是否正確。

# --hints--

你應該輸入 `'Hello'`。確保從行首開始輸入，並且不要忘記引號。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
