---
id: 6ac0a4463662d9738c393e21
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

在前一步，你創建了一個**字串**：一連串被單引號或雙引號包住的字元。在程式設計中，字串用來表述文字。

**終端機**是你可以看到程式碼輸出結果的地方。

為了能看到輸出，將你現有的程式碼行改成：

```py
print('Hello')
```

確保從該行的開頭開始。

# --hints--

你的程式碼中應該有 `print('Hello')`。確保從該行的開頭開始。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
