---
id: 6ac0aaba914d81b32567c076
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

`print` 是一個內建函式（程式），你可以用它將輸出送到終端機。如你所見，文字 `Hello` 現在已經在終端機上顯示。

現在，前往凸顯標示的那一行，輸入文字 `freeCodeCamp` 並列印到終端機：

```py
print('freeCodeCamp')
```

確定從該行的開頭開始輸入。

# --hints--

你的程式碼中應該有 `print('freeCodeCamp')`。確定從該行的開頭開始輸入。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
