---
id: 6ac0ae563f6de70930f47ef7
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

當你撰寫程式碼時，經常會想要重複使用或參考程式碼中已存在的值。變數讓你能做到這一點。

要為變數指定值，請寫下變數名稱，接著是指定運算子（`=`），然後是你想指定給該變數的值：

```py
sentence = 'I am learning Python'
```

前往醒目標示的那一行，創建一個名為 `sentence` 的變數，並將字串 `I am learning Python` 指定給它。

# --hints--

你的程式碼中應該有 `sentence = 'I am learning Python'`。確保從該行的開頭開始寫。

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
