---
id: 660ef55dd468079679ee0092
title: Schritt 12
challengeType: 1
dashedName: step-12
---

# --description--

Sie sind jetzt bereit, Ihre nächste Variable zu deklarieren. Entfernen Sie beide `console.log`-Anweisungen und die `character`-Zuweisung.

Entfernen Sie außerdem Ihre `secondCharacter`-Variable, aber lassen Sie die `character`-Initialisierung unverändert.

# --hints--

Sie sollten keine `console.log()`-Anweisungen in Ihrem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Sie sollten den Wert von `character` nicht neu zuweisen.

```js
assert.isAtMost(__helpers.removeJSComments(code).match(/character\s*=/g).length, 1);
```

Ihre `character`-Variable sollte den Wert `"Hello"` haben.

```js
assert.equal(character, "Hello");
```

Sie sollten keine `secondCharacter`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /secondCharacter/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
secondCharacter = character;
console.log(secondCharacter);
--fcc-editable-region--
```
