---
id: 660ef19b95d3308e7dd31bb6
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

In der <dfn>Konsole</dfn> kann die JavaScript-Ausgabe ausgegeben und angezeigt werden. Mit `console.log()` kannst du Informationen an die Konsole senden. Beispielsweise wird dieser Code `"Naomi"` in der Konsole ausgeben:

```js
let developer = "Naomi";
console.log(developer);
```

Der obige Code greift auf die `developer`-Variable mit dem Namen im `consloe.log()` zu. Beachte, dass der Wert zwischen den Klammern der Wert ist, der ausgegeben wird.

Gibt den Wert der `character`-Variable in der Konsole aus. Klicke dann auf "Console", um die JavaScript Konsole zu sehen.

# --hints--

Die Initialisierung der `character` - Variablen sollte nicht geändert werden.

```js
assert.strictEqual(character, 'Hello');
```

You should access the `console` in your code.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

You should access the `log` property of the `console`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

You should use parentheses to call the `.log()` method.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

You should print the `character` variable to the console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
