---
id: 660ef0f7c4b8e68ccd1f0786
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

JavaScript hat sieben primitive Datentypen, wobei String einer davon ist. In JavaScript stellt ein <dfn>String</dfn> eine Folge von Zeichen dar und kann entweder in einfache (`'`) oder doppelte (`"`) Anführungszeichen eingeschlossen werden.

Beachten Sie, dass Strings <dfn>unveränderlich</dfn> sind, was bedeutet, dass sie nach ihrer Erstellung nicht mehr geändert werden können. Die Variable kann jedoch weiterhin einen anderen Wert zugewiesen bekommen.

Ändern Sie Ihren `"Hello"`-String so, dass einfache Anführungszeichen verwendet werden.

# --hints--

Sie sollten keine doppelten Anführungszeichen in Ihrem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Sie sollten für Ihren `"Hello"`-String einfache Anführungszeichen verwenden.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Sie sollten in Ihrem Code weiterhin `let` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Sie sollten weiterhin eine `character`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Ihre `character`-Variable sollte weiterhin den String `"Hello"` als Wert haben.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
