---
id: 660ef19b95d3308e7dd31bb6
title: Schritt 5
challengeType: 1
dashedName: step-5
---

# --description--

Die <dfn>Konsole</dfn> ermöglicht es Ihnen, JavaScript-Ausgaben auszugeben und anzusehen. Sie können Informationen an die Konsole senden, indem Sie `console.log()` verwenden. Zum Beispiel gibt dieser Code `"Naomi"` in der Konsole aus:

```js
let developer = "Naomi";
console.log(developer);
```

Der obige Code greift auf die `developer`-Variable mit ihrem Namen in den `console.log()` zu. Beachten Sie, dass der Wert zwischen den runden Klammern der Wert ist, der ausgegeben wird.

Geben Sie den Wert der `character`-Variable in der Konsole aus. Klicken Sie dann auf die Schaltfläche „Konsole“, um die JavaScript-Konsole anzusehen.

# --hints--

Die Initialisierung der `character`-Variable sollte nicht geändert werden.

```js
assert.strictEqual(character, 'Hello');
```

Sie sollten in Ihrem Code auf die `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Sie sollten auf die `log`-Eigenschaft der `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Sie sollten die runden Klammern verwenden, um die `.log()`-Methode aufzurufen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Sie sollten die `character`-Variable in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*character\s*\)/);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
