---
id: 660ef5c1904955978a986a5c
title: Schritt 15
challengeType: 1
dashedName: step-15
---

# --description--

Verwenden Sie `let`, um eine `count`-Variable zu deklarieren. Weisen Sie ihr den <dfn>Zahlenwert</dfn> `8` zu. Bei der Verwendung eines Zahlenwerts verwenden Sie keine Anführungszeichen. Zum Beispiel:

```js
let money = 100;
```

# --hints--

Sie sollten ein zweites `let`-Schlüsselwort in Ihrem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

Sie sollten `let` verwenden, um eine `count`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+count/);
```

Sie sollten der `8`-Variable die Zahl `count` zuweisen.

```js
assert.equal(count, 8);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';

--fcc-editable-region--
```
