---
id: 660eebd83100d37862268781
title: Schritt 3
challengeType: 1
dashedName: step-3
---

# --description--

Ihre `character`-Variable hat derzeit keinen Wert. Sie können einen Wert mit dem <dfn>Zuweisungsoperator</dfn> `=` zuweisen. Zum Beispiel:

```js
let hello = "Hello";
```

Einem Wert einer Variable zum Zeitpunkt ihrer Deklaration zuzuweisen, nennt man <dfn>Initialisierung</dfn>.

Initialisieren Sie Ihre `character`-Variable, indem Sie ihr während der Deklaration den Wert `"Hello"` zuweisen.

# --hints--

Sie sollten den Zuweisungsoperator `=` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character\s*=/);
```

Sie sollten den String `"Hello"` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /('|")Hello\1/);
```

Sie sollten für Ihren `"Hello"`-String Anführungszeichen verwenden.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

Sie sollten `"Hello"` Ihrer `character`-Variable zuweisen.

```js
assert.equal(character, "Hello");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character;
--fcc-editable-region--
```
