---
id: 660ef31a5be625914a0102cd
title: Schritt 6
challengeType: 1
dashedName: step-6
---

# --description--

Wenn eine Variable mit dem Schlüsselwort `let` deklariert wird, können Sie den Wert dieser Variable später <dfn>neu zuweisen</dfn> (oder ändern). In diesem Beispiel wird der Wert von `programmer` von `"Naomi"` auf `"CamperChan"` geändert.

```js
let programmer = "Naomi";
programmer = "CamperChan";
```

Beachten Sie, dass Sie beim Neu-Zuweisen einer Variable das Schlüsselwort `let` **nicht** erneut verwenden.

Weisen Sie nach Ihrer `console.log` der Variablen `"World"` den Wert `character` zu.

# --hints--

Sie sollten `character` zweimal in Ihrem Code verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character/g), 3);
```

Sie sollten den Zuweisungsoperator verwenden, um `character` neu zuzuweisen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\s*\=/g), 2);
```

Sie sollten den String `"World"` in Ihrem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /("|')World\1/);
```

Ihre Variable `character` sollte den Wert `"World"` haben.

```js
assert.equal(character, "World");
```

Ihre Neu-Zuweisung sollte `let` nicht verwenden.

```js
assert.isBelow(__helpers.removeJSComments(code).match(/let/g).length, 2);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);

--fcc-editable-region--
```
