---
id: 660ef5105b8ba095307a0e50
title: Schritt 7
challengeType: 1
dashedName: step-7
---

# --description--

Geben Sie nun Ihre Variable `character` erneut in der Konsole aus. Sie sollten zuerst den String `"Hello"` und dann den String `"World"` in der Konsole sehen.

# --hints--

Sie sollten `console.log` ein zweites Mal verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Sie sollten `character` ein zweites Mal in der Konsole ausgeben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*character\s*\)/g), 2);
```


Ihre neue `console.log()` sollte nach Ihrer erneuten Zuweisung kommen.

```js
const lines = __helpers.removeJSComments(code).split(/\n+/)
const reassign = lines.findIndex(l => l.match(/character\s+=\s+("|')World\1/));
const secondLog = lines.findLastIndex(l => l.match(/console\.log/));
assert.isBelow(reassign, secondLog);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";

--fcc-editable-region--
```
