---
id: 660ef857f2806aa626d29d17
title: Schritt 17
challengeType: 1
dashedName: step-17
---

# --description--

Sie können auch Subtraktion (`-`), Multiplikation (`*`) und Division (`/`) durchführen. Experimentieren Sie gerne mit den Operatoren und Zahlen in Ihrem `console.log`. Wenn Sie bereit sind weiterzumachen, entfernen Sie das `console.log`.

# --hints--

Sie sollten keinen `console.log()` in Ihrem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Ihre `character`-Variable sollte weiterhin den Wert `"Hello"` haben.

```js
assert.equal(character, "Hello");
```

Ihre `count`-Variable sollte weiterhin den Wert `8` haben.

```js
assert.equal(count, 8);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
--fcc-editable-region--
let count = 8;
console.log(count + 1);
--fcc-editable-region--
```
