---
id: 660ee6e3a242da6bd579de69
title: Schritt 2
challengeType: 1
dashedName: step-2
---

# --description--

Eines der wichtigsten Konzepte in der Programmierung sind Variablen. Eine <dfn>Variable</dfn> verweist auf eine bestimmte Speicheradresse, die einen Wert speichert. Variablen erhalten einen Namen, mit dem Sie im gesamten Code auf diesen Wert zugreifen können.

Eine Variable zu deklarieren bedeutet, ihr einen Namen zu geben. In JavaScript geschieht dies oft mit dem `let`-Schlüsselwort. Zum Beispiel so deklarieren Sie eine `hello`-Variable:

```js
let hello;
```

Die Benennung von Variablen folgt bestimmten Regeln: Namen können Buchstaben, Zahlen, Dollarzeichen und Unterstriche enthalten, dürfen aber keine Leerzeichen haben und dürfen nicht mit einer Zahl beginnen.

Verwenden Sie das `let`-Schlüsselwort, um eine Variable namens `character` zu deklarieren.

_Hinweis_: Es ist üblich, Anweisungen in JavaScript mit einem Semikolon zu beenden. `;`

# --hints--

Sie sollten `let` in Ihrem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Sie sollten `character` in Ihrem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /character/);
```

Sie sollten `let` verwenden, um eine `character`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Ihre Deklaration sollte mit einem Semikolon enden.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
