---
id: 660ef0f7c4b8e68ccd1f0786
title: Schritt 4
challengeType: 1
dashedName: step-4
---

# --description--

Javascript hat sieben primitive Datentypen. String ist einer davon. In Javascript repräsentiert ein <dfn>string</dfn> eine Zeichenkette von Symbolen und kann entweder mit einzelnen (`'`) oder doppelten Anführungszeichen (`"` umschlossen werden.

Beachte, dass Strings <dfn>immutable</dfn> sind, was bedeutet, dass sie, nachdem sie erstellt wurden, nicht mehr abgeändert werden können. (Die Variable an sich allerdings schon, da du den Wert überschreiben kannst). Der Variable kann trotzdem verändert werden, indem ihr ein anderer Wert zugewiesen wird.

Änder deinen `"Hello"` string so, dass einzelne Anführungszeichen verwendet werden.

# --hints--

Du solltest keine doppelten Anführungszeichen in deinem Code haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /"/);
```

Du solltest einfache Anführungszeichen für den `"Hello"`-Text verwenden.

```js
assert.match(__helpers.removeJSComments(code), /'Hello'/);
```

Du solltest immer noch `let` in deinem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /let/);
```

Du solltest immer noch eine `character`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+character/);
```

Deine `character`-Variable sollte immer noch den String`"Hello"` als Wert haben.

```js
assert.equal(character, "Hello");
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
--fcc-editable-region--
```
