---
id: 660ef6355e8f5a9e67fe5f46
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Mit dem Datentyp `number` können Sie mathematische Operationen wie Addition durchführen. Versuchen Sie, `count + 1` in der Konsole auszugeben.

# --hints--

Sie sollten in Ihrem Code auf `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Sie sollten auf `.log()`s `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Sie sollten `count + 1` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
