---
id: 660ef6355e8f5a9e67fe5f46
title: Schritt 16
challengeType: 1
dashedName: step-16
---

# --description--

Mit dem `number` Datentyp kannst du mathematische Operationen, wie Additionen durchführen. Versuche `count + 1` in der Konsole auszugeben.

# --hints--

Du solltest im Code auf die `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console/);
```

Du solltest auf die `.log()` Eigenschaft der `console` zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Du solltest `count + 1` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*count\s*\+\s*1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
