---
id: 660f1cedf3676fe26122ebf6
title: Schritt 38
challengeType: 1
dashedName: step-38
---

# --description--

Sie sollten die Zahlen von null bis sieben in Ihrer Konsole ausgeben sehen, jeweils eine pro Zeile. Dies dient als Grundlage für die Erstellung Ihrer Pyramide.

Ersetzen Sie Ihre Log-Anweisung durch eine Anweisung, die `i` in Ihr `rows` Array schiebt.

# --hints--

Sie sollten keinen `console.log` Aufruf haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Sie sollten `.push()` auf Ihrem `rows` Array aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Sie sollten `i` in Ihr `rows` Array schieben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Ihr `.push()` sollte in Ihrer `for` Schleife stattfinden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
