---
id: 660f4a83373de83ea101685f
title: Schritt 100
challengeType: 1
dashedName: step-100
---

# --description--

Da Sie in die entgegengesetzte Richtung iterieren, muss Ihre Schleife laufen, solange `i` größer als `0` ist. Sie können dafür den <dfn>größer als</dfn>-Operator `>` verwenden.

Setzen Sie die Bedingung Ihrer Schleife so, dass sie läuft, wenn `i` größer als `0` ist.

# --hints--

Ihre `for`-Schleife sollte laufen, wenn `i` größer als `0` ist.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; false; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
