---
id: 660f505d02b2bd513a1c3468
title: Schritt 111
challengeType: 1
dashedName: step-111
---

# --description--

Ihre Pyramide ist nicht mehr umgekehrt. Das liegt daran, dass Sie neue Zeilen am **Ende** des Arrays hinzufügen.

Aktualisieren Sie den Schleifenkörper, um neue Zeilen am Anfang des Arrays hinzuzufügen.

# --hints--

Sie sollten die `unshift`-Methode von `rows` verwenden.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /\.unshift/);
```

Sie sollten einen `padRow()`-Aufruf als Argument für Ihre `.unshift()`-Methode übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(/);
```

Sie sollten `i` als erstes Argument an Ihren `padRow()`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i/)
```

Sie sollten `count` als zweites Argument an Ihren `padRow()`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.unshift\(\s*padRow\(\s*i\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
