---
id: 660f2eccfe3f820304af1b39
title: Schritt 64
challengeType: 1
dashedName: step-64
---

# --description--

Um zu wissen, wie eine Zeile formatiert werden soll, muss Ihre `padRow`-Funktion wissen, auf welcher Zeilennummer Sie sich befinden und wie viele Zeilen insgesamt erzeugt werden.

Der beste Weg, dies zu erreichen, ist, dafür Funktionsparameter zu erstellen. Geben Sie Ihrer `padRow`-Funktion einen `rowNumber`- und einen `rowCount`-Parameter. Mehrere Parameter werden durch ein Komma getrennt:

```js
function name(first, second) {

}
```

# --hints--

Ihre `padRow`-Funktion sollte einen `rowNumber`-Parameter haben.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber/);
```

Sie sollten nach Ihrem `rowNumber`-Parameter ein Komma hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*/);
```

Ihre `padRow`-Funktion sollte einen `rowCount`-Parameter haben.

```js
assert.match(__helpers.removeJSComments(code), /function\s+padRow\s*\(\s*rowNumber\s*,\s*rowCount\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
