---
id: 660f34626216270c682e2f7b
title: Schritt 66
challengeType: 1
dashedName: step-66
---

# --description--

Ein <dfn>Funktionsaufruf</dfn> ermöglicht es Ihnen, eine Funktion tatsächlich zu verwenden. Vielleicht ist es Ihnen nicht bewusst gewesen, aber die Methoden wie `.push()`, die Sie bisher verwendet haben, waren Funktionsaufrufe.

Eine Funktion wird aufgerufen, indem Sie den Namen der Funktion referenzieren und `()` hinzufügen. So rufen Sie eine `test`-Funktion auf:

```js
test();
```

Ersetzen Sie das `character.repeat(i + 1)` in Ihrem `.push()`-Aufruf durch einen Funktionsaufruf für Ihre `padRow`-Funktion. Fügen Sie noch keine Argumente hinzu.

# --hints--

Sie sollten `i + 1` nicht in Ihrem `push`-Aufruf verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /repeat\(\s*i\s*\+\s*1\s*\)/);
```

Sie sollten `character.repeat` nicht in Ihrem `.push()`-Aufruf verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /push\(\s*character/);
```

Sie sollten `padRow` in Ihrem `.push()`-Aufruf aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\((.+?)?\)\)/);
```

Ihr `padRow`-Aufruf sollte keine Argumente enthalten.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*?padRow\(\s*?\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
