---
id: 660f280dda5040f707c76b4a
title: Schritt 45
challengeType: 1
dashedName: step-45
---

# --description--

Jetzt haben Sie eine Reihe von `#`-Zeichen, aber die Pyramidenform fehlt noch. Glücklicherweise stellt die Variable `i` die aktuelle „Zeilen“-Nummer in Ihrer Schleife dar, sodass Sie sie verwenden können, um eine pyramidenähnliche Struktur zu erstellen.

Um dies zu erreichen, verwenden Sie die `.repeat()`-Methode, die für Strings verfügbar ist. Diese Methode akzeptiert eine Zahl als Argument, die angibt, wie oft der Zielstring wiederholt werden soll. Zum Beispiel können Sie mit `.repeat()` den String `"Code! Code! Code!"` erzeugen:

```js
const activity = "Code! ";
activity.repeat(3);
```

Verwenden Sie die `.repeat()`-Methode für Ihren `character` und geben Sie ihr `i` als Zahl.

# --hints--

Sie sollten die `.repeat()`-Methode verwenden.

```js
assert.match(__helpers.removeJSComments(code), /\.repeat\(/);
```

Sie sollten die `.repeat()`-Methode für Ihre `character`-Variable verwenden.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(/);
```

Sie sollten `i` an Ihre `.repeat()`-Methode übergeben.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\)/)
```

Sie sollten die `.repeat()`-Methode in der `.push()`-Methode verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\.repeat\(\s*i\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
