---
id: 660f4934fb48f63abd5ae371
title: Schritt 95
challengeType: 1
dashedName: step-95
---

# --description--

Die Verwendung von `done`, um die Anzahl der generierten Zeilen zu verfolgen, funktioniert zwar, aber Sie können die Logik tatsächlich noch etwas aufräumen.

Arrays haben eine spezielle `length`-Eigenschaft, mit der Sie sehen können, wie viele Werte oder <dfn>Elemente</dfn> sich im Array befinden. Auf diese Eigenschaft greifen Sie mit einer Syntax wie `myArray.length` zu.

Beachten Sie, dass `rows.length` im `padRow`-Aufruf einen Off-by-one-Fehler verursachen würde, weil `done` *vor* dem Aufruf inkrementiert wird.

Aktualisieren Sie Ihre Bedingung so, dass geprüft wird, ob `rows.length` kleiner als `count` ist.

# --hints--

Ihre `while`-Schleife sollte prüfen, ob `rows.length` kleiner als `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*rows\.length\s*<\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (done <= count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
