---
id: 660f1e3f047bf4e403268713
title: Schritt 39
challengeType: 1
dashedName: step-39
---

# --description--

Unglücklicherweise kannst du jetzt nicht sehen, was deine Schleife macht.

Nutze `let`, um eine `result` Variable zu deklarieren und weise ihr einen leeren String zu. Ein leerer String wird durch Anführungszeichen ohne Zeichen dazwischen dargestellt, wie hier: `""`.

# --hints--

Du solltest eine `result`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Du solltest `let` verwenden, um deine `result`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Deine `result` Variable sollte ein leerer String sein.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
