---
id: 660f1bf673487ae0bb25b900
title: Schritt 37
challengeType: 1
dashedName: step-37
---

# --description--

Deine Schleife sollte nun achtmal laufen. Lass dir innerhalb der Schleife den Wert des `i` Zählers ausgeben und schau dir an, was passiert.

# --hints--

Du solltest `console.log()` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Du solltest den Wert von `i` loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Du solltest den Wert von `i` in deiner `for` Schleife loggen.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
