---
id: 660f0ee51d7460ce88cd248d
title: Schritt 29
challengeType: 1
dashedName: step-29
---

# --description--

Hatten Sie erwartet, `4` in der Konsole zu sehen? `.push()` gibt die neue Länge des Arrays zurück, nachdem Sie den Wert hinzugefügt haben.

Es ist wichtig, sich bewusst zu sein, welche Werte eine Methode zurückgibt. Nehmen Sie sich etwas Zeit, um mit `.push()` und `.pop()` zu experimentieren. Wenn Sie bereit sind, entfernen Sie alle Ihre `.push()`- und `.pop()`-Aufrufe sowie Ihre `console.log`-Anweisungen.

# --hints--

Sie sollten keinen `.push()`-Aufruf haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.push\(/);
```

Sie sollten keinen `.pop()`-Aufruf haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.pop\(/);
```

Sie sollten keine Log-Anweisungen haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log/);
```

Sie sollten keine `popped`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /popped/);
```

Sie sollten keine `pushed`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /pushed/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let pushed = rows.push("freeCodeCamp");
console.log(pushed);
let popped = rows.pop();
console.log(popped);
console.log(rows);
--fcc-editable-region--
```
