---
id: 660f255022991ef34ed0ee88
title: Schritt 44
challengeType: 1
dashedName: step-44
---

# --description--

Das Ausgeben von Zahlen führt nicht zu einer optisch ansprechenden Pyramide. Da Sie nun den formatierten Inhalt Ihres `rows`-Arrays ausgeben, ist es Zeit, Ihre ursprüngliche Schleife zu aktualisieren.

Statt `i` in das Array zu schieben, schieben Sie den Wert Ihrer `character`-Variablen.

# --hints--

Sie sollten Ihre `i`-Variable nicht mehr schieben.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Sie sollten Ihre `character`-Variable schieben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*character\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
