---
id: 660f383d4c772c12ff59904b
title: Schritt 70
challengeType: 1
dashedName: step-70
---

# --description--

Sie können vollständige Ausdrücke als Argument übergeben. Die Funktion erhält das Ergebnis der Auswertung dieses Ausdrucks. Zum Beispiel würden diese beiden Funktionsaufrufe unten dasselbe Ergebnis liefern. (Beachten Sie, dass die Regel der Reihenfolge der Operationen PEMDAS – Klammern, Exponenten, Multiplikation, Division, Addition, Subtraktion – gilt):

```js
test(2 * 3 + 1);
test(7);
```

Wenn Sie sich das Muster noch einmal ansehen:

```md
··#··
·###·
#####
```

Aktualisieren Sie den Wert von `character`, sodass er `2 * rowNumber - 1` Mal wiederholt wird.

Öffnen Sie erneut die Konsole, um das aktualisierte Ergebnis zu sehen.

# --hints--

Ihre Funktion sollte `2 * rowNumber - 1` an Ihren `.repeat()`-Aufruf übergeben.

```js
assert.equal(padRow(4, 5), " ####### ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(rowNumber) + " ".repeat(rowCount - rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
