---
id: 660f53ad3d39175c5d4335ac
title: Schritt 117
challengeType: 1
dashedName: step-117
---

# --description--

Ihr Pyramiden-Generator ist jetzt fertiggestellt und verfügt über mehr Funktionen als ursprünglich geplant! Der nächste Schritt ist, Ihren Code aufzuräumen.

Entfernen Sie alle Kommentare, sowohl einzeilige als auch mehrzeilige, aus Ihrem Code.

# --hints--

Sie sollten keine einzeiligen Kommentare mehr in Ihrem Code haben.

```js
assert.notMatch(code, /\/\//);
```

Sie sollten keine mehrzeiligen Kommentare mehr in Ihrem Code haben.

```js
assert.notMatch(code, /(?:\*\/|\/\*)/);
```

Sie sollten keine Kommentare mehr in Ihrem Code haben.

```js
assert.equal(code, __helpers.removeJSComments(code));
```

Sie sollten Code entfernen, der durch mehrzeilige Kommentare auskommentiert wurde.

```js
assert.notMatch(code, /while/);
assert.lengthOf(code.match(/for\s*\(/g), 2)
assert.lengthOf(code.match(/rows\.push/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
for (let i = 1; i <= count; i++) {
  if (inverted) {
    rows.unshift(padRow(i, count));
  } else {
    rows.push(padRow(i, count));
  }
}

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
