---
id: 660f455b044d3230ed971e98
title: Schritt 85
challengeType: 1
dashedName: step-85
---

# --description--

Eine <dfn>`while`</dfn>-Schleife wird immer wieder ausgeführt, bis die angegebene `condition` nicht mehr wahr ist. Sie hat folgende Syntax:

```js
while (condition) {
  logic;
}
```

Verwenden Sie diese Syntax, um eine `while`-Schleife mit `continueLoop` als Bedingung zu deklarieren. Der Schleifenkörper soll leer sein.

# --hints--

Sie sollten eine `while`-Schleife verwenden.

```js
assert.match(__helpers.removeJSComments(code), /while/);
```

Ihre `while`-Schleife soll `continueLoop` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
