---
id: 660f3b664421471aa595170f
title: Schritt 74
challengeType: 1
dashedName: step-74
---

# --description--

Die Pyramide sieht jetzt etwas seltsam aus. Da Sie die Schleife bei `1` anstatt bei `0` starten, müssen Sie `i` nicht um eins erhöhen, wenn Sie es an `padRow` übergeben.

Aktualisieren Sie das erste Argument Ihres `padRow`-Aufrufs auf `i`.

# --hints--

Sie sollten `i + 1` nicht an `padRow` übergeben.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Sie sollten `i` an `padRow` übergeben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
