---
id: 660f374d532dc41189cc9cc2
title: Schritt 69
challengeType: 1
dashedName: step-69
---

# --description--

Jetzt ist es Zeit für ein bisschen Mathematik. Betrachten Sie eine Pyramide mit drei Reihen. Wenn wir sie zentriert haben wollen, würde sie ungefähr so aussehen:

```md
··#··
·###·
#####
```

Leere Stellen wurden zur besseren Lesbarkeit durch Interpunkts oder Mittelpunkte ersetzt. Wenn Sie das Muster fortsetzen, sehen Sie, dass die Leerzeichen am Anfang und Ende einer Reihe einem Muster folgen.

Aktualisieren Sie Ihre Leerzeichen-Strings so, dass sie `rowCount - rowNumber` mal wiederholt werden.

Öffnen Sie die Konsole, um das Ergebnis zu sehen.

# --hints--

Sie sollten `.repeat()` auf Ihre `" "` Strings aufrufen, um sie `rowCount - rowNumber` mal zu wiederholen.

```js
assert.equal(padRow(1, 3), "  #  ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return " " + character.repeat(rowNumber) + " ";
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
