---
id: 660f34e99571070d56d2f231
title: Schritt 67
challengeType: 1
dashedName: step-67
---

# --description--

Ihre `padRow`-Funktion hat zwei Parameter, die Sie definiert haben. Werte werden diesen Parametern übergeben, wenn eine Funktion aufgerufen wird.

Die Werte, die Sie einem Funktionsaufruf übergeben, werden als <dfn>Argumente</dfn> bezeichnet, und Sie <dfn>übergeben</dfn> Argumente an einen Funktionsaufruf. Hier ist ein Funktionsaufruf mit `"Hello"` als übergebenem Argument:

```js
test("Hello");
```

Übergeben Sie `i + 1` und `count` als Argumente an Ihren `padRow`-Aufruf. Wie bei Parametern werden Argumente durch ein Komma getrennt.

# --hints--

Sie sollten `i + 1` an Ihren `padRow()`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1/);
```

Nach Ihrem `i + 1`-Argument sollte ein Komma stehen.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*/);
```

Sie sollten `count` als Ihr zweites Argument übergeben.

```js
assert.match(__helpers.removeJSComments(code), /push\(\s*padRow\(\s*i\s*\+\s*1\s*,\s*count\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}


--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow())
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
