---
id: 660f3dd626be3a1ffe27e5d1
title: Schritt 77
challengeType: 1
dashedName: step-77
---

# --description--

JavaScript unterstützt auch mehrzeilige Kommentare. Ein mehrzeiliger Kommentar beginnt mit `/*` und endet mit `*/`.

Im Gegensatz zu einem einzeiligen Kommentar umfasst ein mehrzeiliger Kommentar mehrere Zeilen.

Verwenden Sie `/*` und `*/`, um Ihre aktuelle `for`-Schleife einschließlich des Inhalts in einen mehrzeiligen Kommentar zu verwandeln.

# --hints--

Sie sollten einen mehrzeiligen Kommentar mit `/*` beginnen.

```js
assert.match(code, /\/\*/);
```

Sie sollten einen mehrzeiligen Kommentar mit `*/` beenden.

```js
assert.match(code, /\*\//);
```

Ihre gesamte `for`-Schleife sollte auskommentiert werden.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for/g), 1);
assert.notMatch(stripped, /rows\.push/);
assert.notMatch(stripped, /i <= count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
