---
id: 660f5179b3b0ca558f6b4d4f
title: Schritt 113
challengeType: 1
dashedName: step-113
---

# --description--

Verwenden Sie eine `if`-Anweisung, um zu überprüfen, ob `inverted` wahr ist. Denken Sie daran, dass Sie hier keinen Gleichheitsoperator verwenden müssen.

# --hints--

Sie sollten eine `if`-Anweisung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

Ihre `if`-Anweisung sollte prüfen, ob `inverted` `true` ist.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

Ihre `if`-Bedingung sollte keine Vergleichsoperatoren verwenden.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
