---
id: 660f3ce51f70571e1c5227c8
title: Schritt 76
challengeType: 1
dashedName: step-76
---

# --description--

Kommentare können hilfreich sein, um zu erklären, warum Ihr Code einen bestimmten Ansatz verfolgt, oder um To-do-Notizen für Ihr zukünftiges Ich zu hinterlassen.

In JavaScript können Sie `//` verwenden, um einen einzeiligen Kommentar in Ihrem Code zu hinterlassen.

Fügen Sie über Ihrer `for`-Schleife einen einzeiligen Kommentar hinzu, der Sie daran erinnert, den Code in eine andere Art von Schleife zu ändern.

# --hints--

Sie sollten einen einzeiligen Kommentar mit `//` beginnen.

```js
assert.match(code, /\/\//);
```

Ihr einzeiliger Kommentar sollte mindestens fünf Zeichen lang sein.

```js
assert.match(code, /\/\/.{5,}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--

for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
