---
id: 660f359af3e32e0f1a6880b7
title: Schritt 68
challengeType: 1
dashedName: step-68
---

# --description--

Sie sollten jetzt dieselbe Zeichenfolge in Ihrer Konsole sehen. Ihre `padRow`-Funktion macht genau dasselbe wie zuvor, aber jetzt ist sie in einem wiederverwendbaren eigenen Abschnitt.

Verwenden Sie den Addition-Operator, um ein einzelnes Leerzeichen `" "` an den Anfang und das Ende Ihres wiederholten `character`-Strings anzuhängen.

Denken Sie daran, dass Sie den `+`-Operator verwenden können, um Strings so zu verketten:

```js
" " + "string"
```

# --hints--

Sie sollten ein einzelnes Leerzeichen an den Anfang Ihres zurückgegebenen Werts anhängen.

```js
assert.match(padRow(1, 1), /^\s/);
```

Sie sollten ein einzelnes Leerzeichen an das Ende Ihres zurückgegebenen Werts anhängen.

```js
assert.match(padRow(1, 1), /\s$/);
```

Ihre `padRow()`-Funktion sollte die wiederholte `character`-Reihe mit einem Leerzeichen davor und danach zurückgeben.

```js
assert.equal(padRow(1, 1), " # ");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {
  return character.repeat(rowNumber);
}
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
