---
id: 660f207334fabaeac3269c38
title: Schritt 41
challengeType: 1
dashedName: step-41
---

# --description--

Um den `result`-String zu manipulieren, verwenden Sie eine andere Art von Schleife. Genauer gesagt eine `for...of`-Schleife, die jedes Element eines iterierbaren Objekts durchläuft und es vorübergehend einer Variablen zuweist.

Die Syntax für eine `for...of`-Schleife sieht folgendermaßen aus:

```js
for (const value of iterable) {

}
```

Beachten Sie, dass Sie `const` verwenden können, da die Variable nur für eine einzelne Iteration existiert, nicht während der gesamten Schleife.

Erstellen Sie eine `for...of`-Schleife, um Ihr `rows`-Array durchzulaufen und jeden Wert einer `row`-Variablen zuzuweisen.

# --hints--

Sie sollten ein weiteres `for`-Schlüsselwort verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for/g), 2);
```

Sie sollten eine `row`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /\s+row\s+/);
```

Sie sollten `const` verwenden, um Ihre `row`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+row\s+/);
```

Ihre `for...of`-Schleife sollte Ihre `row`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+/);
```

Ihre `row`-Variable sollte mit dem Schlüsselwort `rows` aus `of` extrahiert werden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)/);
```

Ihr `for...of`-Schleifenkörper sollte leer sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*const\s+row\s+of\s+rows\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--

console.log(result);
```
