---
id: 660f47afe4c98536715d5fa4
title: Schritt 91
challengeType: 1
dashedName: step-91
---

# --description--

Der <dfn>strikte Ungleichheitsoperator</dfn> `!==` ermöglicht es Ihnen zu prüfen, ob zwei Werte ungleich sind oder nicht denselben Typ haben. Die Syntax ähnelt der des Gleichheitsoperators: `value !== 4`.

Derzeit läuft die `while`-Schleife nur, wenn `continueLoop` wahr ist. Aktualisieren Sie die `while`-Schleifenbedingung, um zu prüfen, ob `done` ungleich `count` ist.

# --hints--

Ihre `while`-Schleife sollte prüfen, ob `done` und `count` ungleich sind.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*(?:done\s*!==\s*count|count\s*!==\s*done)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;
  rows.push(padRow(done, count));
  if (done === count) {
    continueLoop = false;
  } 
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
