---
id: 660f4efcb8068e4cb470dca1
title: Schritt 109
challengeType: 1
dashedName: step-109
---

# --description--

Nachdem Sie diese Methoden ausprobiert haben, können Sie einen weiteren umgekehrten Pyramidenansatz machen. Aber zuerst müssen Sie Ihre Experimente aufräumen.

Entfernen Sie Ihr `numbers` Array sowie die Methodenaufrufe und Log-Aufrufe.

# --hints--

Sie sollten kein `numbers` Array mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /numbers/);
```

Sie sollten keine `unshifted` Variable mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /unshifted/);
```

Sie sollten keine `shifted` Variable mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /shifted/);
```

Sie sollten Ihre `console.log` Anweisungen nicht mehr haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 1);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
const numbers = [1, 2, 3];
const shifted = numbers.shift();
console.log(shifted);
const unshifted = numbers.unshift(5);
console.log(unshifted);
console.log(numbers);
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
