---
id: 660f2a70ad6225fa503e71c3
title: Schritt 46
challengeType: 1
dashedName: step-46
---

# --description--

Sie kommen dem Ziel näher! An diesem Punkt stoßen Sie auf einen sogenannten <dfn>Off-by-one-Fehler</dfn>, ein häufiges Problem in nullbasierten Indexierungssprachen wie JavaScript.

Der erste Index Ihres `rows`-Arrays ist `0`, weshalb Sie Ihre `for`-Schleife mit `i = 0` starten. Aber das Wiederholen eines Strings nullmal führt dazu, dass nichts ausgegeben wird.

Um das zu beheben, addieren Sie `1` zum Wert von `i` in Ihrem `.repeat()`-Aufruf. Weisen Sie es nicht wie in Ihren Schleifenbedingungen wieder `i` zu.

# --hints--

Sie sollten `1` zu `i` in Ihrer `.repeat()`-Methode hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i))
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
