---
id: 660f1cedf3676fe26122ebf6
title: Schritt 38
challengeType: 1
dashedName: step-38
---

# --description--

Du solltest die Zahlen 1 bis 7 in deiner Konsole sehen (Jeweils eine Zeile für eine Ziffer). Dies ist unsere Grundlage um die Pyramide zu generieren.

Ersetze deine log-Anweisung durch eine push-Anweisung, die `i` in dein `rows`-Array hinzufügt.

# --hints--

Du solltest kein `console.log`-Aufruf mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console/);
```

Führe die `.push()` deines `rows`-Arrays aus.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Du sollst `i` zu deinem `rows`-Array hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*i\s*\)/);
```

Deine `.push()`-Anweisung sollte in der `for`-Schleife ausgeführt werden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*rows\.push\(\s*i\s*\);?\s*\}/)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
  console.log(i);
}
--fcc-editable-region--
```
