---
id: 660f0da9bf1035c9097af20a
title: Schritt 28
challengeType: 1
dashedName: step-28
---

# --description--

Sie sollten `"freeCodeCamp"` in der Konsole ausgegeben gesehen haben. Das liegt daran, dass `.pop()` den Wert zurückgibt, der aus dem Array entfernt wurde – und Sie haben `"freeCodeCamp"` zuvor ans Ende des Arrays angehängt.

Aber was gibt `.push()` zurück? Weisen Sie Ihr bestehendes `rows.push()` einer neuen `pushed`-Variablen zu und geben Sie diese aus.

# --hints--

Sie sollten eine `pushed`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /pushed/);
```

Sie sollten `let` verwenden, um Ihre `pushed`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+pushed/);
```

Sie sollten `rows.push("freeCodeCamp")` Ihrer `pushed`-Variable zuweisen.

```js
assert.equal(pushed, 4);
```

Sie sollten Ihre `pushed`-Variable ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*pushed\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
rows.push("freeCodeCamp");

--fcc-editable-region--
let popped = rows.pop();
console.log(popped);
console.log(rows);
```
