---
id: 660f4b641290da41b2cf0dd9
title: Schritt 103
challengeType: 1
dashedName: step-103
---

# --description--

Genau wie bei der Addition gibt es verschiedene Operatoren, die Sie für die Subtraktion verwenden können. Der <dfn>Subtraktionszuweisungsoperator</dfn> `-=` zieht den angegebenen Wert vom aktuellen Variablenwert ab und weist das Ergebnis dann der Variablen wieder zu.

Ersetzen Sie Ihre Iterationsanweisung durch die korrekte Anweisung mit dem Subtraktionszuweisungsoperator.

# --hints--

Ihre `for` Schleife sollte `i = i - 1` nicht verwenden.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /for\(leti=count;i>0;i=i-1\)\{rows\.push\(padRow\(i,count\)\);/);
```

Ihre `for` Schleife sollte die Subtraktionszuweisung verwenden, um `i` um `1` zu verringern.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*-=\s*1\s*\)\s*\{\s*rows\.push\(\s*padRow\s*\(\s*i\s*,\s*count\s*\)\s*\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; i = i - 1) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
