---
id: 660f415b76859a2736771607
title: Schritt 78
challengeType: 1
dashedName: step-78
---

# --description--

Ihre Pyramide ist wieder verschwunden. Das ist in Ordnung – das ist zu erwarten.

Bevor Sie Ihre neue Schleife erstellen, müssen Sie etwas über `if`-Anweisungen lernen. Eine <dfn>`if`-Anweisung</dfn> ermöglicht es Ihnen, einen Codeblock nur dann auszuführen, wenn eine Bedingung erfüllt ist. Sie verwenden die folgende Syntax:

```js
if (condition) {
  logic
}
```

Erstellen Sie eine `if`-Anweisung mit dem Boolean `true` als Bedingung. Geben Sie im Rumpf den String `"Condition is true"` aus.

# --hints--

Sie sollten eine `if`-Anweisung erstellen.

```js
assert.match(__helpers.removeJSComments(code), /if/);
```

Ihre `if`-Anweisung sollte `true` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)/);
```

Ihr `if`-Rumpf sollte `"Condition is true"` ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*true\s*\)\s*\{\s*console\.log\(\s*("|')Condition is true\1\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
