---
id: 660f033cf051ebb50ea3bf48
title: Schritt 18
challengeType: 1
dashedName: step-18
---

# --description--

In der Programmierung müssen Sie oft mit vielen Daten arbeiten. Es gibt viele Datenstrukturen, die Ihnen helfen können, Ihre Daten zu organisieren und zu verwalten. Eine der grundlegendsten Datenstrukturen ist ein <dfn>Array</dfn>.

Ein <dfn>Array</dfn> ist ein nicht-primitiver Datentyp, der eine Reihe von Werten speichern kann. Nicht-primitive Datentypen unterscheiden sich von primitiven Datentypen dadurch, dass sie komplexere Daten speichern können. Primitive Datentypen wie Strings und Zahlen können jeweils nur einen Wert speichern.

Arrays werden mit eckigen Klammern (`[]`) dargestellt. Hier ist ein Beispiel für eine Variable mit dem Wert eines leeren Arrays:

```js
let array = [];
```

Deklarieren Sie eine `rows`-Variable und weisen Sie ihr ein leeres Array zu.

# --hints--

Sie sollten eine `rows`-Variable haben.

```js
assert.match(__helpers.removeJSComments(code), /rows/);
```

Sie sollten `let` verwenden, um Ihre `rows`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+rows/);
```

Sie sollten Ihrer `rows`-Variable ein leeres Array zuweisen.

```js
assert.deepEqual(rows, []);
```


# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;

--fcc-editable-region--
```
