---
id: 660f1b6e60bd9edf902c81fd
title: Schritt 36
challengeType: 1
dashedName: step-36
---

# --description--

Ihre <dfn>Iteration</dfn>-Anweisung sagt Ihrer Schleife, was sie mit dem Iterator nach jedem Durchlauf machen soll.

Wenn Sie einer Variablen einen neuen Wert zuweisen, können Sie die Variable verwenden, um auf den vorherigen Wert vor der Zuweisung zuzugreifen. So können Sie zum Beispiel drei zu einer bestehenden Zahl addieren. Zum Beispiel würde `bees = bees + 3` den Wert von `bees` um drei erhöhen.

Verwenden Sie diese Syntax, um Ihren `"iteration"`-String durch eine Zuweisung zu ersetzen, die `i` um eins erhöht.

# --hints--

Sie sollten zu Ihrer `i`-Variablen eins addieren.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\s*1/);
```

Sie sollten `i + 1` wieder Ihrer `i`-Variablen zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Ihre `for`-Schleife sollte `i` bei jeder Iteration um `1` erhöhen.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; "iteration") {

}
--fcc-editable-region--
```
