---
id: 660f061d259bbebc37461080
title: Schritt 20
challengeType: 1
dashedName: step-20
---

# --description--

Sie können auf die Werte innerhalb eines Arrays zugreifen, indem Sie den <dfn>Index</dfn> des Werts verwenden. Ein Index ist eine Zahl, die die Position des Werts im Array darstellt, beginnend bei `0` für den ersten Wert.

Sie können den Wert mit der <dfn>Klammernotation</dfn> zugreifen, zum Beispiel `array[0]`.

Verwenden Sie `console.log` und die Klammernotation, um den ersten Wert in Ihrem `rows` Array auszugeben.

# --hints--

Sie sollten eine `console.log()` Anweisung in Ihrem Code haben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/);
```

Sie sollten auf Ihr `rows` Array zugreifen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows/g), 2);
```

Sie sollten die Klammernotation mit Ihrem `rows` Array verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\[/);
```

Sie sollten die Klammernotation verwenden, um auf das erste Element Ihres `rows` Arrays zuzugreifen.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*0\s*\]/);
```

Sie sollten das erste Element Ihres `rows` Arrays ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*]\s*\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];

--fcc-editable-region--
```
