---
id: 660f4990b1caa03b9dc97a43
title: Schritt 97
challengeType: 1
dashedName: step-97
---

# --description--

Sie benötigen Ihre `done`-Variable nicht mehr. Entfernen Sie die Inkrement-Operation aus Ihrer Schleife sowie die Deklaration der `done`-Variable.

# --hints--

Sie sollten die `done`-Variable nicht inkrementieren.

```js
assert.notMatch(__helpers.removeJSComments(code), /done\+\+/);
```

Sie sollten keine `done`-Variable mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /done/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let done = 0;

while (rows.length < count) {
  done++;
  rows.push(padRow(rows.length + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
