---
id: 660f4a1472f8e63d76162ce5
title: Schritt 99
challengeType: 1
dashedName: step-99
---

# --description--

Was wäre, wenn Sie Ihre Pyramide auf den Kopf stellen, also <dfn>invertieren</dfn>? Zeit, es auszuprobieren!

Beginnen Sie damit, eine neue `for`-Schleife zu erstellen. Deklarieren Sie Ihren Iterator `i` und weisen Sie ihm den Wert von `count` zu, dann verwenden Sie den Boolean `false` für Ihre Bedingungs- und Iterationsanweisungen.

# --hints--

Ihr Code sollte eine `for`-Schleife enthalten.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/for\s*\(/g), 2);
```

Ihre `for`-Schleife sollte `i` mit dem Wert von `count` initialisieren.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count/);
```

Ihre `for`-Schleife sollte `false` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false/);
```


Ihre `for`-Schleife sollte `false` als Iteration verwenden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*false\s*;\s*false\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
