---
id: 660f1a00ac619ddc1e259a66
title: Schritt 35
challengeType: 1
dashedName: step-35
---

# --description--

Die <dfn>Bedingung</dfn> einer `for`-Schleife gibt an, wie oft die Schleife iterieren soll. Wenn die `condition` false wird, stoppt die Schleife.

In JavaScript kann ein Boolean-Wert entweder `true` oder `false` sein. Diese sind keine Strings – den Unterschied lernen Sie später noch genauer kennen.

Für den Moment verwenden Sie den <dfn>Kleiner-als-Operator</dfn> (`<`). Damit können Sie prüfen, ob der Wert links kleiner ist als der Wert rechts. Zum Beispiel würde `count < 3` zu `true` ausgewertet werden, wenn `count` `2` ist, und zu `false`, wenn `count` `4` ist.

Ersetzen Sie Ihren `"condition"`-String durch eine Bedingung, die prüft, ob `i` kleiner ist als `count`.

# --hints--

Sie sollten den Kleiner-als-Operator verwenden.

```js
assert.match(__helpers.removeJSComments(code), /</);
```

Sie sollten den Kleiner-als-Operator verwenden, um zu prüfen, ob `i` kleiner ist als `count`.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Ihre `for`-Schleife sollte `i < count` als Bedingung verwenden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0\s*;\s*i\s*<\s*count\s*;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; "condition"; "iteration") {

}
--fcc-editable-region--
```
