---
id: 660f3b664421471aa595170f
title: Step 74
challengeType: 1
dashedName: step-74
---

# --description--

Die Pyramide sieht jetzt ein wenig lustig aus. Weil du die Schleife bei `1` anstatt bei `0`startest, musst du kein zu `i` hinzufügen, wenn du es an `PadRow` weitergibst.

Aktualisier das erste Argument deines `padRow` Aufrufs mit `i`.

# --hints--

Du solltest nicht `i + 1` an `PadRow` übergeben.

```js
assert.notMatch(__helpers.removeJSComments(code), /padRow\(\s*i\s*\+\s*1/);
```

Du solltest `i ` an `padRow` übergeben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*i\s*,/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
