---
id: 660f4f79e2a82a4e92290f44
title: Schritt 110
challengeType: 1
dashedName: step-110
---

# --description--

Manchmal möchten Sie vorherigen Code, den Sie auskommentiert haben, wiederherstellen. Das können Sie tun, indem Sie die `/*` und `*/` um diesen Code entfernen. Das nennt man <dfn>Entkommentieren</dfn>.

Kommentieren Sie nur Ihre erste `for` Schleife aus. Lassen Sie den einzeiligen Kommentar und die anderen beiden mehrzeiligen Kommentare unverändert.

# --hints--

Sie sollten Ihre erste `for` Schleife entkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.match(stripped, /for\s*\(\s*let\s+i\s*=\s*1;\s*i\s*<=\s*count;\s*i\+\+\s*\)/)
```

Sie sollten Ihren einzeiligen Kommentar nicht entfernen.

```js
assert.match(code, /\/\//);
```

Sie sollten Ihre `while` Schleife nicht entkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.notMatch(stripped, /while/);
```

Sie sollten Ihre zweite `for` Schleife nicht entkommentieren.

```js
const stripped = __helpers.removeJSComments(code);
assert.lengthOf(stripped.match(/for\s*\(/g), 2)
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
