---
id: 660f2fbd45b520046cac68e8
title: Schritt 65
challengeType: 1
dashedName: step-65
---

# --description--

Erinnern Sie sich, dass Sie in einem früheren Schritt etwas über Rückgabewerte gelernt haben. Eine Funktion kann einen Wert zurückgeben, den Ihre Anwendung separat verwenden kann.

In einer Funktion wird das Schlüsselwort `return` verwendet, um einen Rückgabewert anzugeben. Zum Beispiel würde diese Funktion den Wert zurückgeben, der dem ersten Parameter übergeben wird:

```js
function name(parameter) {
  return parameter;
}
```

Verwenden Sie das Schlüsselwort `return`, um den Wert der Variablen `character` zurückzugeben, der `rowNumber` mal wiederholt wird.

# --hints--

Sie sollten die Methode `.repeat()` verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/\.repeat\(/g), 2);
```

Sie sollten die Methode `.repeat()` auf Ihrer Variablen `character` verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/character\.repeat\(/g), 2);
```

Sie sollten `rowNumber` an Ihren Aufruf `.repeat()` übergeben.

```js
assert.match(__helpers.removeJSComments(code), /character\.repeat\(\s*rowNumber\s*\)/);
```

Sie sollten das Schlüsselwort `return` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /return/);
```

Sie sollten das Ergebnis Ihres Aufrufs `.repeat()` zurückgeben.

```js
assert.equal(padRow(3), "###");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(rowNumber, rowCount) {

}
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
