---
id: 660f0f980e98e8cf77f1ce31
title: Schritt 30
challengeType: 1
dashedName: step-30
---

# --description--

Ändern Sie Ihre `rows`-Deklaration so, dass ihr wieder ein leeres Array zugewiesen wird.

Ändern Sie außerdem Ihren `'Hello'`-String so, dass er wieder doppelte Anführungszeichen verwendet. Im Allgemeinen ist es egal, welche der beiden Varianten Sie bevorzugen, aber Sie sollten in Ihrem gesamten Projekt konsequent bei dieser Wahl bleiben.

# --hints--

Ihr `rows`-Array sollte leer sein.

```js
assert.empty(rows);
```

Ihr `"Hello"`-String sollte doppelte Anführungszeichen verwenden.

```js
assert.match(__helpers.removeJSComments(code), /"Hello"/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
```
