---
id: 660f07d231941bc11719f664
title: Schritt 21
challengeType: 1
dashedName: step-21
---

# --description--

Arrays sind besonders, weil sie als <dfn>veränderbar</dfn> gelten. Das bedeutet, Sie können den Wert an einem Index direkt ändern.

Zum Beispiel würde dieser Code die Zahl `25` dem zweiten Element im Array zuweisen:

```js
let array = [1, 2, 3];
array[1] = 25;
console.log(array); // prints [1, 25, 3]
```

Aktualisieren Sie das **dritte** Element Ihres `rows`-Arrays auf die Zahl `10`. Geben Sie dann das `rows`-Array in Ihrer Konsole aus.

# --hints--

Sie sollten erneut die Klammernotation für das `rows`-Array verwenden.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/rows\[/g), 2)
```

Sie sollten auf das dritte Element des `rows`-Arrays zugreifen.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]/);
```

Sie sollten den Zuweisungsoperator für das dritte Element des `rows`-Arrays verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\[\s*2\s*\]\s*=/);
```

Sie sollten den Wert `10` dem dritten Element Ihres `rows`-Arrays zuweisen.

```js
assert.equal(rows[2], 10);
```

Sie sollten eine zweite `console.log`-Anweisung in Ihrem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Sie sollten das `rows`-Array ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
console.log(rows[0]);
--fcc-editable-region--

--fcc-editable-region--
```
