---
id: 660f18f059fe0fda192ce394
title: Schritt 34
challengeType: 1
dashedName: step-34
---

# --description--

Ihre Schleife benötigt jetzt einen richtigen Iterator. Der <dfn>Iterator</dfn> ist eine Variable, die Sie speziell in Ihrer `for`-Schleife deklarieren können, um zu steuern, wie die Schleife iteriert oder Ihre Logik durchläuft.

Es ist eine gängige Konvention, `i` als Ihre Iteratorvariable in einer Schleife zu verwenden. Eine `for`-Schleife erlaubt es Ihnen, dies in den runden Klammern `()` zu deklarieren. Zum Beispiel ist hier eine `for`-Schleife, die eine `index`-Variable deklariert und ihr den Wert `100` zuweist.

```js
for (let index = 100; "second"; "third") {

}
```

Ersetzen Sie den String `"iterator"` durch eine `let`-Deklaration für die Variable `i`. Weisen Sie ihr den Wert `0` zu, um zu starten. Dadurch erhält die `i`-Variable beim **ersten Mal**, wenn Ihre Schleife läuft, den Wert `0`.

# --hints--

Sie sollten `let` verwenden, um eine `i`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i/);
```

Sie sollten `0` Ihrer `i`-Variable zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /let\s+i\s*=\s*0/);
```

Ihre `for`-Schleife sollte einen `i`-Iterator bei `0` starten.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for ("iterator"; "condition"; "iteration") {

}
--fcc-editable-region--
```
