---
id: 660f20473aef47e9b8c9afc6
title: Schritt 40
challengeType: 1
dashedName: step-40
---

# --description--

Fügen Sie eine Log-Anweisung hinzu, um den Wert von `result` auszugeben. Je nachdem, welche Konsole Sie verwenden, sehen Sie möglicherweise keine Ausgabe.

# --hints--

Sie sollten einen `console.log`-Aufruf hinzufügen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(/);
```

Sie sollten Ihre `result`-Variable ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--

--fcc-editable-region--
```
