---
id: 660f45ccf4ca5c31f253005a
title: Schritt 86
challengeType: 1
dashedName: step-86
---

# --description--

Im Moment, wenn Sie `continueLoop` auf true ändern, läuft Ihre `while`-Schleife unendlich weiter. Das nennt man eine <dfn>Endlosschleife</dfn>, und Sie sollten vorsichtig sein, diese zu vermeiden. Eine Endlosschleife kann Ihr System blockieren, sodass ein kompletter Neustart nötig ist, um wieder herauszukommen.

Um dies zu vermeiden, verwenden Sie zunächst den Inkrementoperator, um den Wert der `done`-Variablen innerhalb Ihrer Schleife zu erhöhen.

# --hints--

Ihre `while`-Schleife sollte die `done`-Variable inkrementieren.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
