---
id: 660f0a55847d6cc485f29ba5
title: Schritt 26
challengeType: 1
dashedName: step-26
---

# --description--

Eine <dfn>Methode</dfn> in JavaScript ist eine Funktion, die mit bestimmten Werten oder Objekten verknüpft ist. Ein Beispiel, das Sie bereits kennengelernt haben, ist die `.log()`-Methode, die Teil des `console`-Objekts ist.

Arrays haben ihre eigenen Methoden, und die erste, die Sie erkunden werden, ist die `.push()`-Methode. Damit können Sie einen Wert ans Ende eines Arrays „pushen“. Hier ist ein Beispiel, um die Zahl `12` ans Ende eines Arrays hinzuzufügen:

```js
array.push(12);
```

Verwenden Sie `.push()`, um den String `"freeCodeCamp"` ans Ende Ihres `rows`-Arrays hinzuzufügen. Fügen Sie diesen Code vor Ihrem `console.log` ein, damit Sie die Änderung an Ihrem Array sehen können.

# --hints--

Sie sollten `.push()` in Ihrem Code verwenden.

```js
assert.match(__helpers.removeJSComments(code), /\.push\(/);
```

Sie sollten die `.push()`-Methode Ihres `rows`-Arrays verwenden.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(/);
```

Sie sollten den String `"freeCodeCamp"` an Ihre `.push()`-Methode übergeben.

```js
assert.match(__helpers.removeJSComments(code), /rows\.push\(\s*('|")freeCodeCamp\1\s*\);?/)
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];

console.log(rows);
--fcc-editable-region--
```
