---
id: 660f46460f9c36330ebc07d8
title: Schritt 88
challengeType: 1
dashedName: step-88
---

# --description--

Der Gleichheitsoperator kann in JavaScript zu einigen merkwürdigen Verhaltensweisen führen. Zum Beispiel ist `"0" == 0` wahr, obwohl einer ein String und der andere eine Zahl ist.

Der <dfn>strikte Gleichheitsoperator</dfn> `===` wird verwendet, um zu prüfen, ob zwei Werte gleich sind und denselben Typ haben. Als allgemeine Regel sollten Sie immer diesen Gleichheitsoperator verwenden. Mit dem strikten Gleichheitsoperator wird `"0" === 0` falsch, weil sie zwar denselben Wert null haben könnten, aber nicht denselben Typ.

Aktualisieren Sie Ihre `done == count`-Bedingung, damit sie den strikten Gleichheitsoperator verwendet.

# --hints--

Ihre `if`-Bedingung sollte strikte Gleichheit verwenden.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*===\s*count|count\s*===\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--
  if (done == count) {

  }
--fcc-editable-region--
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
