---
id: 660f17d4e9f227d86e834abd
title: Schritt 33
challengeType: 1
dashedName: step-33
---

# --description--

Um eine Pyramide zu erzeugen, müssen Sie mehrere Reihen erstellen. Wenn Sie eine Aufgabe wiederholt ausführen müssen, bis eine Bedingung erfüllt ist, verwenden Sie eine <dfn>Schleife</dfn>. Es gibt viele Möglichkeiten, eine Schleife zu schreiben.

Sie beginnen mit einer einfachen `for`-Schleife. `for`-Schleifen verwenden die folgende Syntax:

```js
for (iterator; condition; iteration) {
  logic;
}
```

In den kommenden Schritten werden Sie jede Komponente einer Schleife im Detail untersuchen. Erstellen Sie vorerst eine `for`-Schleife, die die Begriffe `"iterator"`, `"condition"` und `"iteration"` für die drei Komponenten enthält. Lassen Sie den <dfn>Rumpf</dfn> der Schleife, den Abschnitt innerhalb der geschweiften Klammern `{}`, leer.


# --hints--

Sie sollten eine `for`-Schleife haben.

```js
assert.match(__helpers.removeJSComments(code), /for/);
```

Die erste Komponente Ihrer `for`-Schleife sollte der String `"iterator"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1/);
```

Die zweite Komponente Ihrer `for`-Schleife sollte der String `"condition"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2/);
```

Die dritte Komponente Ihrer `for`-Schleife sollte der String `"iteration"` sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)/);
```

Der Rumpf Ihrer `for`-Schleife sollte leer sein.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*('|")iterator\1\s*;\s*('|")condition\2\s*;\s*('|")iteration\3\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--
```
