---
id: 660f3ba3cceef11b6ba08b59
title: Schritt 75
challengeType: 1
dashedName: step-75
---

# --description--

Leider ist jetzt die Unterseite der Pyramide verschwunden. Das liegt daran, dass Sie einen weiteren <dfn>Off-by-one-Fehler</dfn> gemacht haben.

Ihre ursprüngliche Schleife lief für `i` Werte von `0` bis `7`, weil `count` `8` ist und Ihre Bedingung verlangt, dass `i` kleiner als `count` ist. Ihre Schleife läuft jetzt für `i` Werte von `1` bis `7`.

Ihre Schleife muss so aktualisiert werden, dass sie auch läuft, wenn `i` `8` ist. Wenn man sich Ihre Logik anschaut, bedeutet das, dass Ihre Schleife laufen sollte, wenn `i` <dfn>kleiner gleich</dfn> `count` ist. Dafür können Sie den Kleiner-gleich-Operator `<=` verwenden.

Aktualisieren Sie Ihre Schleifenbedingung so, dass sie läuft, solange `i` kleiner gleich `count` ist.

# --hints--

Ihre `for` Schleife sollte nicht prüfen, ob `i` kleiner als `count` ist.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*<\s*count/);
```

Ihre `for` Schleife sollte prüfen, ob `i` kleiner gleich `count` ist.

```js
assert.match(__helpers.removeJSComments(code), /i\s*<=\s*count/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 1; i < count; i++) {
  rows.push(padRow(i, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
