---
id: 660f39b444fd6f16d1e49c1f
title: Schritt 73
challengeType: 1
dashedName: step-73
---

# --description--

Anstatt `i + 1` an Ihren `padRow`-Aufruf übergeben zu müssen, könnten Sie Ihre Schleife stattdessen bei `1` starten. So können Sie eine einsbasierte Schleife erstellen.

Aktualisieren Sie Ihren Iterator, damit er bei `1` statt bei `0` beginnt.

# --hints--

Ihre `for`-Schleife sollte `i` bei `1` initialisieren.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i++) {
  rows.push(padRow(i + 1, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
