---
id: 660f2b6fd54ac1fc142804dd
title: Schritt 47
challengeType: 1
dashedName: step-47
---

# --description--

Die Logik zum Formatieren dieser Pyramide wird wahrscheinlich kompliziert, weshalb es ein guter Zeitpunkt ist, diesen Code in eine Funktion auszulagern.

Eine <dfn>Funktion</dfn> ist ein Codeblock, der in Ihrer Anwendung mehrfach verwendet werden kann. Funktionen werden mit folgender Syntax deklariert:

```js
function name(parameter) {

}
```

Das Schlüsselwort `function` teilt JavaScript mit, dass die Variable `name` eine Funktion sein wird. `parameter` ist eine Variable, die einen Wert darstellt, der beim Aufruf der Funktion übergeben wird. Eine Funktion kann beliebig viele oder wenige <dfn>Parameter</dfn> haben. Wie bei einer `for`-Schleife ist der Bereich zwischen den geschweiften Klammern der <dfn>Funktionskörper</dfn>.

Deklarieren Sie eine `padRow`-Funktion. Erstellen Sie noch keine Parameter-Variablen. Der Funktionskörper soll leer sein. Denken Sie daran, für Ihre Benennungskonvention camel case zu verwenden.

# --hints--

Sie sollten das Schlüsselwort `function` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /function/);
```

Sie sollten eine `padRow`-Funktion deklarieren.

```js
assert.isFunction(padRow);
```

Ihre `padRow()`-Funktion sollte keine Parameter haben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)/);
```

Ihre `padRow()`-Funktion sollte einen leeren Körper haben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\s*\(\s*\)\s*\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
