---
id: 660f4cffb1459d45e34902d1
title: Schritt 106
challengeType: 1
dashedName: step-106
---

# --description--

Sie können die umgekehrte Pyramide tatsächlich bauen, ohne „rückwärts“ zu schleifen, wie Sie es zuvor getan haben.

Dafür müssen Sie ein paar neue Array-Methoden lernen. Beginnen Sie damit, `const` zu verwenden, um eine `numbers`-Variable zu deklarieren. Weisen Sie ihr ein Array mit den Elementen `1`, `2` und `3` zu. Loggen Sie dann das `numbers`-Array.

# --hints--

Sie sollten `const` verwenden, um eine `numbers`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+numbers/);
```

Ihre `numbers`-Variable sollte ein Array sein.

```js
assert.isArray(numbers);
```

Ihr `numbers`-Array sollte die Elemente `1`, `2` und `3` in genau dieser Reihenfolge enthalten.

```js
assert.deepEqual(numbers, [1,2,3]);
```

Sie sollten Ihr `numbers`-Array loggen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*numbers\s*\);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--

--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
