---
id: 660f165270622fd4ec0da3f7
title: Schritt 31
challengeType: 1
dashedName: step-31
---

# --description--

Eine Variable mit dem Schlüsselwort `let` zu deklarieren, erlaubt es, ihr später einen neuen Wert zuzuweisen. Das bedeutet, Sie könnten `character` später einen völlig anderen Wert geben.

Für dieses Projekt möchten Sie diese Variablenwerte jedoch nicht ändern. Stattdessen sollten Sie `const` verwenden, um sie zu deklarieren. `const`-Variablen sind etwas Besonderes.

Erstens kann eine `const`-Variable nicht wie eine `let`-Variable neu zugewiesen werden. Dieser Code würde einen Fehler werfen:

```js
const firstName = "Naomi";
firstName = "Jessica";
```

Eine `const`-Variable darf auch nicht uninitialisiert bleiben. Dieser Code würde einen Fehler werfen:

```js
const firstName;
```

Ersetzen Sie Ihre `let`-Schlüsselwörter durch `const`.

# --hints--

Sie sollten `const` verwenden, um Ihre `character`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+character/);
```

Sie sollten `const` verwenden, um Ihre `count`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+count/);
```

Sie sollten `const` verwenden, um Ihre `rows`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+rows/);
```

Sie sollten `let` nicht in Ihrem Code verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /let/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = "Hello";
let count = 8;
let rows = [];
--fcc-editable-region--
```
