---
id: 660f447efc0e722f016c1be0
title: Schritt 81
challengeType: 1
dashedName: step-81
---

# --description--

Der Text ist wieder aufgetaucht! Das liegt daran, dass `"false"` ein String ist, der beim Auswerten als Boolean zu `true` wird. Das bedeutet, `"false"` ist ein truthy-Wert.

Ein <dfn>truthy-Wert</dfn> ist ein Wert, der beim Auswerten als Boolean als wahr gilt. Die meisten Werte, denen Sie in JavaScript begegnen, sind truthy.

Ein <dfn>falsy-Wert</dfn> ist das Gegenteil – ein Wert, der beim Auswerten als Boolean als falsch gilt. JavaScript hat eine definierte Liste von falsy-Werten. Einige davon sind `false`, `0`, `""`, `null`, `undefined` und `NaN`.

Versuchen Sie, Ihre `if`-Bedingung in einen leeren String `""` zu ändern, der ein falsy-Wert ist.

# --hints--

Ihre `if`-Anweisung sollte `""` als Bedingung haben.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*("|')\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("false") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
