---
id: 660f1e3f047bf4e403268713
title: Schritt 39
challengeType: 1
dashedName: step-39
---

# --description--

Leider können Sie jetzt nicht sehen, was Ihre Schleife macht.

Verwenden Sie `let`, um eine `result`-Variable zu deklarieren, und weisen Sie ihr einen leeren String zu. Ein leerer String wird durch Anführungszeichen dargestellt, zwischen denen nichts steht, zum Beispiel `""`.

# --hints--

Sie sollten eine `result`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /result/);
```

Sie sollten `let` verwenden, um Ihre `result`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+result/);
```

Ihre `result`-Variable sollte ein leerer String sein.

```js
assert.equal(result, "");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

--fcc-editable-region--

--fcc-editable-region--
```
