---
id: 660f1bf673487ae0bb25b900
title: Schritt 37
challengeType: 1
dashedName: step-37
---

# --description--

Ihre Schleife sollte jetzt achtmal ausgeführt werden. Geben Sie im Schleifenkörper den Wert des `i`-Iterators aus und beobachten Sie, was passiert.

# --hints--

Sie sollten `console.log()` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /console\.log/)
```

Sie sollten den Wert von `i` ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*i\s*\)/);
```

Sie sollten den Wert von `i` in Ihrer `for`-Schleife ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*0;\s*i\s*<\s*count;\s*i\s*=\s*i\s*\+\s*1\s*\)\s*\{\s*console\.log\(\s*i\s*\);?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {

}
--fcc-editable-region--
```
