---
id: 660f48a419b40238e2b8b4d5
title: Schritt 93
challengeType: 1
dashedName: step-93
---

# --description--

Deine Schleife hängt nicht mehr von der `-continueLoop`Variablen ab. Dies macht die Variable zu einer <dfn>ungenutzten Deklaration</dfn>. Generell wollen Sie ungenutzte Erklärungen vermeiden, um zukünftiger Verwirrung vorzubeugen.

Entferne deine `ContinueLoop` Variable.

# --hints--

Du solltest keine `continueLoop`-Variable mehr haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /continueLoop/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
let continueLoop = false;
let done = 0;

while (done !== count) {
  done++;
  rows.push(padRow(done, count));
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
