---
id: 660f039ff313dbb696b007ca
title: Schritt 19
challengeType: 1
dashedName: step-19
---

# --description--

Wenn ein Array Werte oder <dfn>Elemente</dfn> enthält, sind diese Werte durch Kommas getrennt. Hier ist ein Array, das zwei Strings enthält:

```js
let array = ["first", "second"];
```

Ändern Sie Ihre `rows`-Deklaration so, dass sie ein Array mit den Strings `"Naomi"`, `"Quincy"` und `"CamperChan"` ist. Die Reihenfolge der Werte in einem Array ist wichtig, also halten Sie sich an diese Reihenfolge. Denken Sie daran, dass Strings Groß- und Kleinschreibung unterscheiden.

# --hints--

Das erste Element in Ihrem Array sollte der String `"Naomi"` sein.

```js
assert.equal(rows[0], "Naomi");
```

Das zweite Element in Ihrem Array sollte der String `"Quincy"` sein.

```js
assert.equal(rows[1], "Quincy");
```

Das dritte Element in Ihrem Array sollte der String `"CamperChan"` sein.

```js
assert.equal(rows[2], "CamperChan");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
let count = 8;
let rows = [];
--fcc-editable-region--
```
