---
id: 660f5179b3b0ca558f6b4d4f
title: Schritt 113
challengeType: 1
dashedName: step-113
---

# --description--

Verwende eine `if` Anweisung, um zu prüfen, ob `inverted` wahr ist. Denk daran, dass Du hier keinen Gleichheits-Operator benötigst.

# --hints--

Du solltest eine `if`-Anweisung nutzen.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(/);
```

Deine `if`-Anweisung sollte prüfen ob `inverted` `true` ist.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted/);
```

Deine `if` Bedingung soll keine Vergleichsoperatoren verwenden.

```js
assert.match(__helpers.removeJSComments(code), /if\s*\(\s*inverted\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];
let inverted = true;

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
--fcc-editable-region--
for (let i = 1; i <= count; i++) {

  rows.unshift(padRow(i, count));
}
--fcc-editable-region--

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

/*for (let i = count; i > 0; i--) {
  rows.push(padRow(i, count));
}*/

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
