---
id: 660f0c34aad72dc712b97624
title: Schritt 27
challengeType: 1
dashedName: step-27
---

# --description--

Eine weitere für dieses Projekt wichtige Methode ist die `.pop()`-Methode. Sie entfernt das letzte Element aus einem Array und <dfn>gibt</dfn> dieses Element zurück.

Wenn eine Methode einen Wert zurückgibt, können Sie sich das so vorstellen, dass sie Ihnen den Wert zurückgibt und ihn für die Verwendung in anderen Teilen Ihres Codes verfügbar macht.

Erstellen Sie eine neue Variable namens `popped` und weisen Sie ihr das Ergebnis von `rows.pop()` zu. Geben Sie dann `popped` in der Konsole aus.

# --hints--

Sie sollten eine Variable namens `popped` deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /popped/);
```

Sie sollten `let` verwenden, um Ihre Variable namens `popped` zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+popped/);
```

Sie sollten die `.pop()`-Methode aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /\.pop\(\s*\)/);
```

Sie sollten die `.pop()`-Methode auf Ihrem `rows`-Array aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /rows\.pop\(\s*\)/);
```

Sie sollten Ihre Variable `popped` in der Konsole ausgeben.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*popped\s*\)/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
--fcc-editable-region--
let rows = ["Naomi", "Quincy", "CamperChan"];
rows.push("freeCodeCamp");

console.log(rows);
--fcc-editable-region--
```
