---
id: 660f3915b41a441537ec9f5e
title: Schritt 72
challengeType: 1
dashedName: step-72
---

# --description--

Da Sie `i` nur um `1` erhöhen, können Sie den <dfn>Inkrementoperator</dfn> `++` verwenden. Dieser Operator erhöht den Wert einer Variablen um 1 und aktualisiert die Zuweisung für diese Variable. Zum Beispiel würde `test` hier zu `8` werden:

```js
let test = 7;
test++;
```

Ersetzen Sie Ihre Additionszuweisung durch den Inkrementoperator für Ihre Schleifeniteration.

# --hints--

Ihre `for`-Schleife sollte keine Additionszuweisung mit `i` verwenden.

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

Ihre `for`-Schleife sollte den Inkrementoperator auf `i` anwenden.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+\+/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i += 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
