---
id: 660f4ae5b3924c3fc3373973
title: Schritt 101
challengeType: 1
dashedName: step-101
---

# --description--

Ihre Iterationsanweisung wird ebenfalls anders sein. Statt bei jeder Schleife `1` zu `i` hinzuzufügen, müssen Sie `1` subtrahieren.

Wie Sie es zuvor mit `i = i + 1` gemacht haben, aktualisieren Sie Ihre Iterationsanweisung so, dass `i` den Wert erhält, der sich daraus ergibt, `1` von sich selbst zu subtrahieren.

# --hints--

Ihre `for`-Schleife sollte `i = i - 1` als Iteration verwenden.

```js
assert.match(__helpers.removeJSComments(code), /for\s*\(\s*let\s+i\s*=\s*count\s*;\s*i\s*>\s*0\s*;\s*i\s*=\s*i\s*-\s*1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

/*while (rows.length < count) {
  rows.push(padRow(rows.length + 1, count));
}*/

--fcc-editable-region--
for (let i = count; i > 0; false) {

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
