---
id: 660f23b53db70af0f2620e78
title: Schritt 43
challengeType: 1
dashedName: step-43
---

# --description--

Jetzt erscheinen alle Ihre Zahlen in derselben Zeile. Das funktioniert nicht, um eine Pyramide zu erstellen.

Sie müssen jeder Zeile eine neue Zeile hinzufügen. Wenn Sie jedoch die Eingabetaste drücken, um in JavaScript einen Zeilenumbruch zwischen Anführungszeichen einzufügen, führt dies zu einem Parsing-Fehler. Stattdessen müssen Sie die spezielle <dfn>Escape-Sequenz</dfn> `\n` verwenden, die beim Protokollieren des Strings als neue Zeile interpretiert wird. Zum Beispiel:

```js
lineOne = lineOne + "\n" + lineTwo;
```

Verwenden Sie einen zweiten Additionsoperator, um nach den Werten `result` und `row` eine neue Zeile anzufügen.

# --hints--

Sie sollten die `\n` Escape-Sequenz verwenden. Denken Sie daran, dass sie ein String sein muss, also in Anführungszeichen steht.

```js
assert.match(__helpers.removeJSComments(code), /('|")\\n\1/);
```

Sie sollten Ihre `row` Variable an Ihre `result` Variable anhängen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row\s*\+\s*('|")\\n\1/);
```

Sie sollten die `\n` Escape-Sequenz an Ihre `row` Variable anhängen.

```js
assert.match(__helpers.removeJSComments(code), /row\s*\+\s*('|")\\n\1/);
```

Sie sollten die gesamte Verkettung wieder Ihrer `result` Variable zuweisen.

```js
assert.strictEqual(result, "0\n1\n2\n3\n4\n5\n6\n7\n");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {
  result = result + row;
}
--fcc-editable-region--

console.log(result);
```
