---
id: 660f17294346b7d69e79db3d
title: Schritt 32
challengeType: 1
dashedName: step-32
---

# --description--

Sie sind jetzt bereit, Ihren Pyramiden-Generator zu erstellen. Ihre `character`-Variable dient als Baustein für die Pyramide.

`"Hello"` funktioniert dafür möglicherweise nicht sehr gut. Ändern Sie den Wert von `character` in das Rautezeichen `"#"`.

# --hints--

Ihre `character`-Variable sollte ein Rautezeichen sein.

```js
assert.equal(character, "#");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const character = "Hello";
const count = 8;
const rows = [];
--fcc-editable-region--
```
