---
id: 660f38c34a4de6141c0c369f
title: Schritt 71
challengeType: 1
dashedName: step-71
---

# --description--

Ihr Pyramiden-Generator funktioniert jetzt wie erwartet. Dies ist jedoch eine ausgezeichnete Gelegenheit, den von Ihnen geschriebenen Code noch genauer zu untersuchen.

Der Additionsoperator ist nicht die einzige Möglichkeit, Werte zu einer Variablen hinzuzufügen. Der <dfn>Additionszuweisungsoperator</dfn> kann als Kurzform verwendet werden und bedeutet: „Nehmen Sie den ursprünglichen Wert der Variablen, addieren Sie diesen Wert und weisen Sie das Ergebnis der Variablen wieder zu.“ Zum Beispiel würden diese beiden Anweisungen dasselbe Ergebnis liefern:

```js
test = test + 1;
test += 1;
```

Aktualisieren Sie Ihre Iterationsanweisung in der `for`-Schleife, damit sie die Additionszuweisung verwendet.

# --hints--

Ihre `for`-Schleife sollte `i = i + 1` nicht verwenden;

```js
assert.notMatch(__helpers.removeJSComments(code), /i\s*=\s*i\s*\+\s*1/);
```

Ihre `for`-Schleife sollte die Additionszuweisung mit `i` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /i\s*\+=\s*1/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

--fcc-editable-region--
for (let i = 0; i < count; i = i + 1) {
--fcc-editable-region--
  rows.push(padRow(i + 1, count));
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
