---
id: 660f09a2694b59c3a10ee304
title: Schritt 23
challengeType: 1
dashedName: step-23
---

# --description--

Entfernen Sie vorerst Ihre erste Konsolenausgabe und Ihre `rows[rows.length - 1]`-Zuweisung. Lassen Sie die zweite `rows`-Ausgabe für später stehen.

# --hints--

Sie sollten Ihre `console.log(rows[0])`-Anweisung entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*rows\[\s*0\s*\]\s*\)/);
```

Sie sollten Ihre `rows[rows.length - 1]`-Neuzuweisung entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /rows\[\s*rows\.length\s*-\s*1\s*\]/);
```

Sie sollten Ihre `console.log(rows)`-Anweisung nicht entfernen.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*rows\s*\);/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
console.log(rows[0]);
rows[rows.length - 1] = 10;
--fcc-editable-region--
console.log(rows);
```
