---
id: 660f4774e3e0df35a68bb5f2
title: Schritt 90
challengeType: 1
dashedName: step-90
---

# --description--

Damit Ihre Pyramide wieder erzeugt wird, fügen Sie das Ergebnis des Aufrufs von `padRow` mit `done` und `count` als Argumente zu Ihrem `rows`-Array hinzu, ähnlich wie Sie es in Ihrer ersten Schleife gemacht haben.

# --hints--

Ihre Schleife sollte die `.push()`-Methode auf Ihrem `rows` aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(/);
```

Sie sollten Ihre `padRow`-Funktion in Ihrer `.push()`-Methode aufrufen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(/);
```

Sie sollten `done` als erstes Argument an Ihren `padRow`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done/);
```

Sie sollten `count` als zweites Argument an Ihren `padRow`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*rows\.push\(\s*padRow\(\s*done\s*,\s*count\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

while (continueLoop) {
  done++;
--fcc-editable-region--

--fcc-editable-region--
  if (done === count) {
    continueLoop = false;
  } 
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
