---
id: 660f229d2dbe09ef2954a4a1
title: Schritt 42
challengeType: 1
dashedName: step-42
---

# --description--

Erinnern Sie sich an Ihre vorherige Schleife, in der Sie den Addition-Operator `+` verwendet haben, um den Wert von `i` um `1` zu erhöhen.

Sie können etwas Ähnliches mit einem String-Wert machen, indem Sie einen neuen String an einen bestehenden String anhängen. Zum Beispiel würde `hello = hello + " World";` den String `" World"` an den bereits gespeicherten String in der Variablen `hello` anhängen. Dies nennt man <dfn>Verkettung</dfn>.

Verwenden Sie in Ihrer `for...of` Schleife den Addition-Operator, um den Wert `row` an den Wert `result` anzuhängen.

# --hints--

Sie sollten den Verkettungsoperator auf Ihre `result` Variable anwenden.

```js
assert.match(__helpers.removeJSComments(code), /(?:result\s*\+|\+\s*result)/);
```

Sie sollten `row` an Ihre `result` Variable anhängen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*\+\s*row|row\s*\+\s*result/);
```

Sie sollten das Ergebnis Ihrer Verkettung wieder der Variablen `result` zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*(result\s*\+\s*row|row\s*\+\s*result);?/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

for (let i = 0; i < count; i = i + 1) {
  rows.push(i);
}

let result = ""

--fcc-editable-region--
for (const row of rows) {

}
--fcc-editable-region--

console.log(result);
```
