---
id: 6610c77d50636722e5b6be17
title: Schritt 60
challengeType: 1
dashedName: step-60
---

# --description--

Unter der `return`-Anweisung protokollieren Sie den String `"This works!"` in der Konsole.

Nachdem Sie das getan haben, werden Sie sehen, dass der String `"This works!"` nicht in der Konsole angezeigt wird und die `console.log("This works!")`-Zeile ausgegraut ist.

Kopieren Sie das Konsolenprotokoll und fügen Sie es über der `return`-Anweisung ein. Nun sollte der String `"This works!"` in der Konsole erscheinen.

Eine wichtige Sache, die Sie über das Schlüsselwort `return` wissen sollten, ist, dass es nicht nur einen Wert definiert, der von Ihrer Funktion zurückgegeben wird, sondern auch die Ausführung Ihres Codes innerhalb einer Funktion oder eines Block-Statements stoppt. Das bedeutet, dass jeder Code nach einer `return`-Anweisung nicht ausgeführt wird.

# --hints--

Ihre `padRow`-Funktion sollte die `test`-Variable zurückgeben.

```js
assert.equal(padRow("Naomi"), "Testing");
```

Ihr erstes `console.log` sollte nach Ihrem `return`-Schlüsselwort kommen.

```js
assert.match(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

Ihr zweites `console.log` sollte vor Ihrem `return`-Schlüsselwort kommen.

```js
assert.match(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  
  return test;

}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
