---
id: 6610c21b3ef82015573ffbbe
title: Schritt 51
challengeType: 1
dashedName: step-51
---

# --description--

Ihre `call`-Variable hat einen `undefined`-Wert, obwohl Sie sie definiert haben! Das liegt daran, dass Ihre `padRow`-Funktion derzeit keinen Wert zurückgibt. Standardmäßig geben Funktionen `undefined` als Wert zurück.

Um etwas anderes zurückzugeben, müssen Sie das Schlüsselwort `return` verwenden. Hier ist ein Beispiel für eine Funktion, die den String `"Functions are cool!"` zurückgibt:

```js
function demo() {
  return "Functions are cool!";
}
```

Verwenden Sie das Schlüsselwort `return`, damit Ihre Funktion den String `"Hello!"` zurückgibt.

# --hints--

Ihre `padRow`-Funktion sollte das Schlüsselwort `return` verwenden.

```js
assert.match(padRow.toString(), /return/);
```

Ihre `padRow`-Funktion sollte den String `"Hello!"` zurückgeben.

```js
assert.equal(padRow(), "Hello!");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
