---
id: 6610c105bbdacc114d6cdc44
title: Schritt 48
challengeType: 1
dashedName: step-48
---

# --description--

Um eine Funktion zu verwenden, müssen Sie sie aufrufen. Ein <dfn>Funktionsaufruf</dfn> weist Ihre Anwendung an, den Code der Funktion an der Stelle auszuführen, an der Sie sie aufrufen. Die Syntax für einen Funktionsaufruf ist der Funktionsname, gefolgt von runden Klammern. Zum Beispiel definiert und ruft dieser Code eine `test` Funktion auf.

```js
function test() {

}
test();
```

Rufen Sie Ihre `padRow` Funktion auf.

# --hints--

Sie sollten die `padRow` Funktion aufrufen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Der Körper Ihrer `padRow` Funktion sollte leer sein.

```js
assert.match(padRow.toString(), /\{\s*\}/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}

--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
