---
id: 6610c538372aa61cc0f5b122
title: Schritt 57
challengeType: 1
dashedName: step-57
---

# --description--

Variablen in JavaScript sind in einem bestimmten <dfn>Gültigkeitsbereich</dfn> verfügbar. Anders gesagt: Wo eine Variable deklariert wird, bestimmt, wo im Code sie verwendet werden kann.

Der erste Gültigkeitsbereich ist der globale Gültigkeitsbereich. Variablen, die außerhalb eines „Blocks“ wie einer Funktion oder `for` Schleife deklariert werden, befinden sich im <dfn>globalen Gültigkeitsbereich</dfn>. Ihre `character`, `count` und `rows` Variablen sind alle im globalen Gültigkeitsbereich.

Wenn eine Variable im globalen Gültigkeitsbereich ist, kann eine Funktion in ihrer Definition darauf zugreifen. Hier ist ein Beispiel für eine Funktion, die eine globale `title` Variable verwendet:

```js
const title = "Professor ";
function demo(name) {
  return title + name;
}
demo("Naomi")
```

Dieses Beispiel würde `"Professor Naomi"` zurückgeben. Aktualisieren Sie Ihre `padRow` Funktion so, dass sie den Wert zurückgibt, der entsteht, wenn Ihre `character` Variable an den Anfang des `name` Parameters angehängt wird.

# --hints--

Ihre `padRow` Funktion sollte `character` an den Anfang von `name` anhängen.

```js
assert.match(padRow.toString(), /character\s*\+\s*name/);
```

Ihre `padRow` Funktion sollte das Ergebnis von `character + name` zurückgeben.

```js
assert.equal(padRow("Naomi"), "#Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
