---
id: 6610c87eac0f0b256d7b037e
title: Schritt 62
challengeType: 1
dashedName: step-62
---

# --description--

Da Ihre Funktion den Parameter nicht mehr verwendet hat, hatte eine Änderung des Arguments keine Auswirkung.

Entfernen Sie nun die `test`-Deklaration und die `return`-Anweisung aus Ihrer `padRow`-Funktion, sodass die Funktion wieder leer ist.

# --hints--

Ihre `padRow`-Funktion sollte keine `test`-Variable enthalten.

```js
assert.notMatch(padRow.toString(), /test/);
```

Ihre `padRow`-Funktion sollte keinen Wert zurückgeben.

```js
assert.notMatch(padRow.toString(), /return/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  const test = "Testing";
  return test;
}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
