---
id: 6610b8017d1671f2814e8c77
title: Schritt 9
challengeType: 1
dashedName: step-9
---

# --description--

Wenn Sie eine Variable deklarieren, ohne sie zu initialisieren, gilt sie als <dfn>nicht initialisiert</dfn>. Derzeit ist Ihre `secondCharacter`-Variable nicht initialisiert.

Fügen Sie eine `console.log()` hinzu, um zu sehen, welchen Wert Ihre `secondCharacter`-Variable hat.

# --hints--

Sie sollten eine zweite `console.log` hinzufügen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log/g), 2);
```

Sie sollten Ihre `secondCharacter`-Variable protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*secondCharacter\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

--fcc-editable-region--
```
