---
id: 6610b9f7619764fad5fd516d
title: Schritt 11
challengeType: 1
dashedName: step-11
---

# --description--

Sie können auch den Wert einer Variablen einer anderen Variablen zuweisen. Zum Beispiel:

```js
let first = "One";
let second = "Two";
second = first;
```

Die `second`-Variable hätte jetzt den Wert `"One"`.

Um das in Aktion zu sehen, ändern Sie Ihre `secondCharacter`-Zuweisung von `"Test"` zu Ihrer `character`-Variablen.

Öffnen Sie dann die Konsole, um zu sehen, was ausgegeben wird.

# --hints--

Sie sollten den Wert `"Test"` nicht Ihrer `secondCharacter`-Variablen zuweisen.

```js
assert.notEqual(secondCharacter, "Test");
```

Sie sollten den Wert der `character`-Variablen Ihrer `secondCharacter`-Variablen zuweisen. Vergessen Sie nicht das Semikolon.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=\s*character;/);
```

Ihre `secondCharacter`-Variable sollte jetzt den Wert `"World"` haben.


```js
assert.equal(secondCharacter, "World");
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;
--fcc-editable-region--
secondCharacter = "Test";
--fcc-editable-region--
console.log(secondCharacter);
```
