---
id: 6610b741b54b90f0c0fb3d58
title: Schritt 8
challengeType: 1
dashedName: step-8
---

# --description--

Wenn Variablennamen aus mehr als einem Wort bestehen, gibt es bestimmte Benennungskonventionen, wie Sie die Wörter großschreiben. In JavaScript ist die Konvention, die Sie verwenden sollten, <dfn>camel case</dfn>.

Camel case bedeutet, dass das erste Wort im Namen komplett in Kleinbuchstaben geschrieben wird, die folgenden Wörter jedoch alle mit Großbuchstaben beginnen. Hier sind einige Beispiele für camel case:

```js
let variableOne;
let secondVariable;
let yetAnotherVariable;
let thisIsAnAbsurdlyLongName;
```

Verwenden Sie camel case, um eine neue `secondCharacter`-Variable zu deklarieren.

# --hints--

Sie sollten eine `secondCharacter`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter/);
```

Sie sollten `let` verwenden, um Ihre `secondCharacter`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter/);
```

Sie sollten Ihrer `secondCharacter`-Variable keinen Wert zuweisen. Vergessen Sie nicht das Semikolon.

```js
assert.match(__helpers.removeJSComments(code), /let\s+secondCharacter;/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
console.log(character);

--fcc-editable-region--
```
