---
id: 6610c2d8d67563174fcf96dc
title: Schritt 52
challengeType: 1
dashedName: step-52
---

# --description--

Wenn Sie einen Wert haben, der explizit in Ihrem Code geschrieben steht, wie der `"Hello!"`-String in Ihrer Funktion, wird dieser als <dfn>hartkodiert</dfn> betrachtet. Einen Wert innerhalb einer Funktion hartzukodieren, macht sie möglicherweise nicht so wiederverwendbar, wie Sie es sich wünschen.

Stattdessen können Sie <dfn>Parameter</dfn> für die Funktion definieren. Parameter sind spezielle Variablen, denen beim Aufrufen der Funktion ein Wert zugewiesen wird und die in Ihrer Funktion verwendet werden können, um das Ergebnis des Codes dynamisch zu verändern.

Um Ihrer Funktion einen Parameter hinzuzufügen, müssen Sie einen Variablennamen innerhalb der runden Klammern einfügen. Zum Beispiel hat diese `demo`-Funktion einen `name`-Parameter:

```js
function demo(name) {

}
```

`name` klingt nach einem nützlichen Parameter, also fügen Sie ihn Ihrer `padRow`-Funktion hinzu.

# --hints--

Ihre `padRow`-Funktion sollte einen `name`-Parameter haben.

```js
assert.match(padRow.toString(), /\(name\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
