---
id: 6610bf6fa14d700beed1b109
title: Schritt 87
challengeType: 1
dashedName: step-87
---

# --description--

Der <dfn>Gleichheitsoperator</dfn> `==` wird verwendet, um zu prüfen, ob zwei Werte gleich sind. Um zwei Werte zu vergleichen, verwenden Sie eine Anweisung wie `value == 8`.

Fügen Sie unter `done++` innerhalb Ihrer Schleife eine `if`-Anweisung hinzu. Die Anweisung soll prüfen, ob `done` gleich `count` ist, und dabei den Gleichheitsoperator verwenden.

# --hints--

Sie sollten in Ihrer Schleife eine `if`-Anweisung verwenden. Diese sollte nach `done++` eingefügt werden.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if/);
```

Ihre `if`-Anweisung sollte den Gleichheitsoperator verwenden, um `done` und `count` in der Bedingung zu vergleichen.

```js
assert.match(__helpers.removeJSComments(code), /while\s*\(\s*continueLoop\s*\)\s*\{\s*done\+\+;\s*if\s*\(\s*(?:done\s*==\s*count|count\s*==\s*done)\s*\)\s*\{/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

let continueLoop = false;
let done = 0;

--fcc-editable-region--
while (continueLoop) {
  done++;

}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
