---
id: 6610c16c4fa0df12c0e30675
title: Schritt 49
challengeType: 1
dashedName: step-49
---

# --description--

Sie rufen Ihre `padRow`-Funktion auf, verwenden aber das Ergebnis dieses Funktionsaufrufs nicht weiter. Alle Funktionen in JavaScript <dfn>geben</dfn> einen Wert zurück, das heißt, sie liefern das definierte Ergebnis ihres Aufrufs, das Sie an anderer Stelle verwenden können.

Um das Ergebnis des Aufrufs Ihrer `padRow`-Funktion zu sehen, deklarieren Sie eine `call`-Variable und weisen Sie dieser den Wert Ihres bestehenden `padRow`-Aufrufs zu.

# --hints--

Sie sollten eine `call`-Variable deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /(const|let|var)\s+call/);
```

Sie sollten `const` verwenden, um Ihre `call`-Variable zu deklarieren.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call/);
```

Sie sollten `call` das Ergebnis Ihres `padRow`-Aufrufs zuweisen.

```js
assert.match(__helpers.removeJSComments(code), /const\s+call\s*=\s*padRow\(\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow() {

}
--fcc-editable-region--
padRow();
--fcc-editable-region--


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
