---
id: 6610c48c4ea0891afa7c4696
title: Schritt 54
challengeType: 1
dashedName: step-54
---

# --description--

Wenn Sie Ihre Konsole erneut öffnen, sehen Sie, dass Ihre `padRow`-Funktion `undefined` zurückgibt, obwohl Sie einen Rückgabewert definiert haben! Das liegt daran, dass Parametern beim **Aufrufen** der Funktion ein Wert übergeben werden muss.

Wenn Sie einer Funktionsaufruf einen Wert übergeben, wird dieser Wert als <dfn>Argument</dfn> bezeichnet. Hier ist ein Beispiel für das Aufrufen einer `demo`-Funktion und das Übergeben von `"Naomi"` als Argument für den `name`-Parameter.

```js
function demo(name) {
  return name;
}
demo("Naomi");
```

Geben Sie Ihren eigenen Namen als Argument für den `name`-Parameter in Ihrem `padRow`-Aufruf an. Denken Sie daran, dass Ihr Name ein String ist, also müssen Sie Anführungszeichen verwenden.

# --hints--

Sie sollten einen String an Ihren `padRow()`-Aufruf übergeben.

```js
assert.match(__helpers.removeJSComments(code), /padRow\(\s*("|').+\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
const call = padRow();
--fcc-editable-region--
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
