---
id: 6610c83b52583e245a079217
title: Schritt 61
challengeType: 1
dashedName: step-61
---

# --description--

Jetzt hat Ihre `call`-Variable den Wert `"Testing"`. Aber Ihre Funktion verwendet den `name`-Parameter nicht mehr.

Entfernen Sie den `name`-Parameter aus Ihrer Funktionsdeklaration, und entfernen Sie dann Ihren `"CamperChan"`-String aus dem `padRow`-Aufruf.

Entfernen Sie außerdem beide `console.log` aus der `padRow`-Funktion.
# --hints--

Ihre `padRow`-Funktion sollte keinen `name`-Parameter haben.

```js
assert.notMatch(padRow.toString(), /name/);
```

Sie sollten keinen `"CamperChan"` an Ihren `padRow`-Aufruf übergeben.

```js
assert.notMatch(__helpers.removeJSComments(code), /CamperChan/);
```

Sie sollten Ihre `padRow`-Funktion weiterhin aufrufen.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/padRow\(\)/g), 2);
```

Sie sollten vor Ihrem `console.log`-Schlüsselwort keine `return` haben.

```js
assert.notMatch(padRow.toString(), /console.log\s*\(\s*('|"|`)This\s+works!\1\);\s+return\s+test;/);
```

Sie sollten nach Ihrem `console.log`-Schlüsselwort keine `return` haben.

```js
assert.notMatch(padRow.toString(), /return\s+test;\s+console.log\s*\(\s*('|"|`)This\s+works!\1\);/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";
  console.log("This works!");
  return test;
  console.log("This works!");
}
const call = padRow("CamperChan");
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
