---
id: 6610c8cfe4cf4d278e35c156
title: Schritt 63
challengeType: 1
dashedName: step-63
---

# --description--

Wie erwartet gibt Ihre Funktion jetzt wieder `undefined` zurück. Ihre `call`-Variable wird nicht mehr benötigt, also entfernen Sie die `call`-Deklaration und das `console.log` für die `call`-Variable.

# --hints--

Sie sollten keine `call`-Deklaration haben.

```js
assert.notMatch(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)), /constcall/);
```

Sie sollten Ihre `call`-Variable nicht protokollieren.

```js
assert.notMatch(__helpers.removeJSComments(code), /call/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow() {

}
const call = padRow();
console.log(call);
--fcc-editable-region--

for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
