---
id: 6610c71600966a2191d3a64a
title: Schritt 59
challengeType: 1
dashedName: step-59
---

# --description--

Werte, die aus einer Funktion zurückgegeben werden, werden durch das Aufrufen der Funktion verwendet. Sie können den Funktionsaufruf direkt als den Wert verwenden, den er zurückgibt, oder den zurückgegebenen Wert in einer Variablen speichern. So können Sie den Wert, der einer lokal gültigen Variablen zugewiesen wurde, außerhalb der Funktion verwenden, in der er erstellt wurde.

```js
function getName() {
  const name = "Camper cat";
  return name;
}

console.log(getName()); // "Camper cat"

const capturedReturnValue = getName();
console.log(capturedReturnValue); // "Camper cat"

console.log(name); // reference error
```

Um Ihren `"Testing"`-Wert zu verwenden, geben Sie ihn aus der `padRow`-Funktion zurück, indem Sie Ihre `return`-Anweisung so aktualisieren, dass nur die `test`-Variable zurückgegeben wird.


# --hints--

Ihre `padRow`-Funktion sollte die `test`-Variable zurückgeben.

```js
assert.equal(padRow("Naomi"), "Testing");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  const test = "Testing";  
  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
