---
id: 6610b8f6a98d25f4d485a94d
title: Schritt 10
challengeType: 1
dashedName: step-10
---

# --description--

Der Standardwert einer nicht initialisierten Variable ist `undefined`. Dies ist ein spezieller Datentyp, der einen Wert darstellt, der noch keine Definition hat.

Sie können einer nicht initialisierten Variable trotzdem einen Wert zuweisen. Hier ist ein Beispiel:

```js
let uninitialized;
uninitialized = "assigned";
```

Weisen Sie den String `"Test"` Ihrer `secondCharacter`-Variable unterhalb Ihrer Deklaration zu. Öffnen Sie die Konsole, um zu sehen, wie sich Ihr Log geändert hat.

# --hints--

Sie sollten `secondCharacter` nicht initialisieren. Denken Sie daran, dass Initialisierung bedeutet, beim Deklarieren der Variable einen Wert zuzuweisen.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+secondCharacter\s*=/);
```

Sie sollten den Zuweisungsoperator bei `secondCharacter` verwenden.

```js
assert.match(__helpers.removeJSComments(code), /secondCharacter\s*=/);
```

Sie sollten den String `"Test"` Ihrer `secondCharacter`-Variable zuweisen.

```js
assert.equal(secondCharacter, "Test");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
let character = 'Hello';
console.log(character);
character = "World";
let secondCharacter;

console.log(secondCharacter);
--fcc-editable-region--
```
