---
id: 6610c6541c82551f95e765ab
title: Schritt 58
challengeType: 1
dashedName: step-58
---

# --description--

Variablen können auch innerhalb einer Funktion deklariert werden. Diese Variablen gelten als im <dfn>lokalen Gültigkeitsbereich</dfn> oder <dfn>Block-Gültigkeitsbereich</dfn> befindlich. Eine innerhalb einer Funktion deklarierte Variable kann nur innerhalb dieser Funktion verwendet werden. Wenn Sie versuchen, außerhalb der Funktion darauf zuzugreifen, erhalten Sie einen Referenzfehler.

Um dies in Aktion zu sehen, verwenden Sie `const`, um eine `test`-Variable in Ihrer `padRow`-Funktion zu deklarieren. Initialisieren Sie sie mit dem Wert `"Testing"`.

Versuchen Sie dann unterhalb Ihrer Funktion, `test` in der Konsole auszugeben. Sie werden einen Fehler sehen, weil sie außerhalb des lokalen Gültigkeitsbereichs der Funktion nicht definiert ist. Entfernen Sie diese `console.log`, um die Tests zu bestehen und fortzufahren.

# --hints--

Ihre Funktion sollte eine `test`-Variable deklarieren.

```js
assert.match(padRow.toString(), /var\s+test/);
```

Sie sollten `test` mit dem Wert `"Testing"` initialisieren. Vergessen Sie nicht das Semikolon.

```js
assert.match(padRow.toString(), /var\s+test\s*=\s*('|")Testing\1;/)
```

Ihre `test`-Variable sollte vor Ihrem `return`-Schlüsselwort stehen.

```js
const str = padRow.toString();
const testIndex = str.indexOf("test");
const returnLocation = str.indexOf("return");
const testLocation = testIndex !== -1 ? testIndex : returnLocation
assert.isBelow(testLocation, returnLocation);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {

  return character + name;
}

--fcc-editable-region--
const call = padRow("CamperChan");
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
