---
id: 6610c424b7119919b62932f4
title: Schritt 53
challengeType: 1
dashedName: step-53
---

# --description--

Eine Funktion muss keinen fest codierten Wert zurückgeben. Sie kann den Wert zurückgeben, der in einer Variablen gespeichert ist. Parameter sind spezielle Variablen für eine Funktion, daher können sie ebenfalls zurückgegeben werden.

Ändern Sie Ihre `padRow`-Funktion so, dass sie den `return`-Parameter direkt `name`.

# --hints--

Ihre `padRow`-Funktion sollte den Wert des `name`-Parameters zurückgeben.

```js
assert.equal(padRow("Naomi"), "Naomi");
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

--fcc-editable-region--
function padRow(name) {
  return "Hello!";
}
--fcc-editable-region--
const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
