---
id: 66458f0a05df478aa627629e
title: Schritt 82
challengeType: 1
dashedName: step-82
---

# --description--

Der Text ist wieder weg! Leere Strings werden als `false` ausgewertet, was einen <dfn>falsy-Wert</dfn> darstellt. Sie werden in zukünftigen Projekten mehr über truthy- und falsy-Werte lernen.

Neben `if`-Anweisungen gibt es in JavaScript auch <dfn>else if</dfn>-Anweisungen. `else if`-Anweisungen ermöglichen es Ihnen, mehrere Bedingungen in einem einzigen Codeblock zu überprüfen.

Hier ist die Syntax für eine `else if`-Anweisung:

```js
if (condition1) {
  // code to run if condition1 is true
} else if (condition2) {
  // code to run if condition2 is true
} else if (condition3) {
  // code to run if condition3 is true
} 
```

Wenn die erste Bedingung `false` ist, prüft JavaScript die nächste Bedingung in der Kette. Wenn die zweite Bedingung `false` ist, prüft JavaScript die dritte Bedingung, und so weiter.

Fügen Sie unter Ihrer `if`-Anweisung eine `else if`-Anweisung hinzu, die überprüft, ob `5` kleiner als `10` ist. Dann protokollieren Sie im Körper der `else if`-Anweisung den String `"5 is less than 10"` in der Konsole.

Überprüfen Sie die Konsole, um die Ergebnisse zu sehen.

# --hints--

Sie sollten eine `else if`-Anweisung haben.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(/);
```

Ihre `else if`-Anweisung sollte prüfen, ob `5` kleiner als `10` ist.

```js
assert.match(__helpers.removeJSComments(code), /else\s+if\s*\(\s*5\s*<\s*10\s*\)/);
```

Sie sollten den String `"5 is less than 10"` in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)5\s+is\s+less\s+than\s+10\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

