---
id: 664599653fcd6e97104f9261
title: Schritt 83
challengeType: 1
dashedName: step-83
---

# --description--

Manchmal möchten Sie anderen Code ausführen, wenn alle `if...else if` Bedingungen `false` sind. Das können Sie tun, indem Sie einen `else` Block hinzufügen.

Ein `else` Block wird nur ausgewertet, wenn die Bedingungen in den `if` und `else if` Blöcken nicht erfüllt sind.

Hier wird der `else` Block zum `else if` Block hinzugefügt.

```js

if (condition) {
  // this code will run if condition is true
} else if (condition2) {
  // this code will run if the first condition is false
} else {
  // this code will run 
  // if the first and second conditions are false
}
```

Fügen Sie einen `else` Block zum `else if` Block hinzu. Innerhalb des `else` Blocks protokollieren Sie den String `"This is the else block"` in der Konsole.

Um die Ergebnisse in der Konsole zu sehen, können Sie manuell den `<` im `else if` Ausdruck zu `>` ändern. Dadurch wird die Bedingung `false` und der `else` Block ausgeführt.

# --hints--

Sie sollten einen `else` Block haben.

```js
assert.match(__helpers.removeJSComments(code), /else\s*\{/);
```

Ihr `else` Block sollte den String `"This is the else block"` in der Konsole protokollieren.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*('|"|`)This\s+is\s+the\s+else\s+block\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(rowNumber, rowCount) {
  return " ".repeat(rowCount - rowNumber) + character.repeat(2 * rowNumber - 1) + " ".repeat(rowCount - rowNumber);
}

// TODO: use a different type of loop
/*for (let i = 1; i <= count; i++) {
  rows.push(padRow(i, count));
}*/

--fcc-editable-region--
if ("") {
  console.log("Condition is true");
} else if (5 < 10) {
  console.log("5 is less than 10");
}
--fcc-editable-region--

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```

