---
id: 664ed97c55a99e5ffab759dc
title: Schritt 13
challengeType: 1
dashedName: step-13
---

# --description--

Bevor Sie weitermachen, sollten Sie einen Moment nutzen, um die Konzepte zu überprüfen, die Sie gelernt haben.

Verwenden Sie das Schlüsselwort `let`, um eine `profession`-Variable und eine `age`-Variable zu deklarieren. Initialisieren Sie `profession` mit dem String `"teacher"`, aber initialisieren Sie `age` nicht mit einem Wert.

Geben Sie beide Variablen in der Konsole aus, um die Ergebnisse zu sehen.

# --hints--

Sie sollten eine `profession`-Variable deklarieren.

```js
assert.match(code, /(?:var|let|const)\s+profession/);
```

Sie sollten `let` verwenden, um die `profession`-Variable zu deklarieren.

```js
assert.match(code, /let\s+profession/);
```

Sie sollten der `"teacher"`-Variable den String `profession` zuweisen.

```js
assert.strictEqual(profession, "teacher");
```

Sie sollten eine `age`-Variable deklarieren.

```js
assert.match(code, /(?:var|let|const)\s+age/);
```

Sie sollten `let` verwenden, um die `age`-Variable zu deklarieren.

```js
assert.match(code, /let\s+age/);
```

Sie sollten der `age`-Variable keinen Wert zuweisen.

```js
assert.isUndefined(age);
```

Sie sollten `age` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*age\s*\);?/);
```

Sie sollten `profession` in der Konsole ausgeben.

```js
assert.match(code, /console\.log\(\s*profession\s*\);?/);
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--

--fcc-editable-region--
```
