---
id: 664ed9344ba3495e1f67f8ea
title: Schritt 14
challengeType: 1
dashedName: step-14
---

# --description--

Nachdem Sie die Deklaration und Initialisierung überprüft haben, entfernen Sie den Code, den Sie für Ihre Überprüfung geschrieben haben.

Entfernen Sie nicht Ihre `character`-Variable.

# --hints--

Sie sollten keine `profession`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+profession/);
```

Sie sollten keine `age`-Variable haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /let\s+age/);
```

Sie sollten keine `console.log`-Anweisung für `age` haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*age\s*\);/);
```

Sie sollten keine `console.log`-Anweisung für `profession` haben.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\(\s*profession\s*\);/);
```

Sie sollten weiterhin Ihre `character`-Variable haben.

```js
assert.strictEqual(character, "Hello");
```

# --seed--

## --seed-contents--

```js
let character = "Hello";
--fcc-editable-region--
let profession = "teacher";
let age;
console.log(age);
console.log(profession);
--fcc-editable-region--
```
