---
id: 665732e62eb8b30c249b5c2c
title: Schritt 25
challengeType: 1
dashedName: step-25
---

# --description--

Jetzt sind Sie bereit, mit dem nächsten Satz von Array-Lektionen fortzufahren.

Entfernen Sie den gesamten Code aus dem vorherigen Schritt.

# --hints--

Sie sollten alle Verweise auf Ihre `cities`-Variable entfernen.

```js
assert.notMatch(__helpers.removeJSComments(code), /cities/);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--
let cities = ["London", "New York", "Mumbai"];
console.log(cities);
cities[cities.length - 1] = "Mexico City";
console.log(cities);
--fcc-editable-region--
console.log(rows);
```
