---
id: 665732da4815b70bb083915e
title: Schritt 24
challengeType: 1
dashedName: step-24
---


# --description--

In den letzten Schritten haben Sie alles über das Arbeiten mit Arrays gelernt. Nehmen Sie sich einen Moment Zeit, um zu überprüfen, was Sie gelernt haben.

Beginnen Sie damit, eine `cities`-Variable zu deklarieren und sie als Array der Strings `"London"`, `"New York"` und `"Mumbai"` zu initialisieren. Geben Sie dann diese Variable in der Konsole aus.

Ändern Sie nach der Ausgabe das letzte Element von `cities` in den String `"Mexico City"` und geben Sie die `cities`-Variable erneut in der Konsole aus.

Wenn alles richtig gemacht wurde, sollten Sie diese Ausgabe in der Konsole sehen.

```js
[ "London", "New York", "Mumbai" ]
[ "London", "New York", "Mexico City" ]
```

# --hints--

Sie sollten `let` verwenden, um eine `cities`-Variable zu deklarieren.

```js
assert.match(code, /let\s+cities/);
```

Sie sollten der `"London"`-Variable ein Array der Strings `"New York"`, `"Mumbai"` und `cities` zuweisen.

```js
assert.match(code, /let\s+cities\s*=\s*\[\s*("|'|`)London\1\s*,\s*("|'|`)New York\2\s*,\s*("|'|`)Mumbai\3\s*,?\s*\]/);
```

Sie sollten `console.log()` verwenden, um das gesamte `cities`-Array in der Konsole auszugeben.

```js
assert.match(code, /console\.log\(\s*cities\s*\)/);
```

Sie sollten das letzte Element des `cities`-Arrays in den String `"Mexico City"` ändern. Denken Sie daran, dass Sie auf das letzte Element eines Arrays mit `array[array.length - 1]` zugreifen können.

```js
assert.deepEqual(cities, ["London", "New York", "Mexico City"]);
```

Sie sollten zwei `console.log(cities)`-Anweisungen in Ihrem Code haben.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log\(\s*cities\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
let character = 'Hello';
let count = 8;
let rows = ["Naomi", "Quincy", "CamperChan"];
--fcc-editable-region--

--fcc-editable-region--
console.log(rows);
```
