---
id: 66643c9a9e0ad1c787f5fe2b
title: Schritt 56
challengeType: 1
dashedName: step-56
---

# --description--

Nachdem Sie diese schnelle Überprüfung abgeschlossen haben, sollten Sie Ihre `addTwoNumbers`-Funktion, `sum`-Variable und die Protokollausgabe entfernen.

# --hints--

Sie sollten keine `addTwoNumbers`-Funktion in Ihrem Code haben.

```js
assert.notMatch(code, /addTwoNumbers/);
```

Sie sollten keine `sum`-Variable in Ihrem Code haben.

```js
assert.notMatch(code, /sum/);
```

Sie sollten die `sum`-Variable nicht protokollieren.

```js
assert.notMatch(code, /console\.log\(\s*sum\s*\)/);
```

# --seed--

## --seed-contents--

```js
const character = "#";
const count = 8;
const rows = [];

function padRow(name) {
  return name;
}
--fcc-editable-region--
function addTwoNumbers(num1, num2) {
  return num1 + num2;
}

const sum = addTwoNumbers(5, 10);
console.log(sum)

--fcc-editable-region--

const call = padRow();
console.log(call);


for (let i = 0; i < count; i = i + 1) {
  rows.push(character.repeat(i + 1))
}

let result = ""

for (const row of rows) {
  result = result + row + "\n";
}

console.log(result);
```
