---
id: 661899113600ef307d8a7487
title: Passaggio 23
challengeType: 0
dashedName: step-23
---

# --description--

Per ottenere la mediana di un array con un numero dispari di elementi, devi trovare e restituire il numero centrale.

Ecco come trovare il numero centrale di un array con un numero dispari di elementi:

```js
arr[Math.floor(arr.length / 2)];
```

Ecco un esempio più lungo che trova il numero centrale di un array con `5` elementi:

```js
const numbers = [1, 2, 3, 4, 5];
const middleNumber = numbers[Math.floor(numbers.length / 2)];
console.log(middleNumber); // 3
```

Il motivo per cui usi `Math.floor` è che il risultato della divisione di un numero dispari per `2` sarà un numero decimale. `Math.floor` arrotonderà per difetto al numero intero più vicino.

Dichiara una variabile `oddListMedian` e assegnale il risultato del trovare il numero centrale dell'array `testArr1`. Poi stampa la variabile `oddListMedian` nella console.

Apri la console per vedere il risultato.

# --hints--

Dovresti avere una variabile `oddListMedian`.

```js
assert.isDefined(oddListMedian);
```

Dovresti assegnare il risultato del trovare il numero centrale per l'array `testArr1` alla variabile `oddListMedian`. Ricorda di fare riferimento all'esempio precedente: `arr[Math.floor(arr.length / 2)];`

```js
assert.strictEqual(oddListMedian, 3);
```

Dovresti avere un'istruzione console che stampa la variabile `oddListMedian`.

```js
assert.match(code, /console\.log\(\s*oddListMedian\s*\);?/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1" />
    <link rel="stylesheet" href="./styles.css" />
    <script src="./script.js"></script>
    <title>Statistics Calculator</title>
  </head>
  <body>
    <h1>Statistics Calculator</h1>
    <p>Enter a list of comma-separated numbers.</p>
    <form onsubmit="calculate(); return false;">
      <label for="numbers">Numbers:</label>
      <input type="text" name="numbers" id="numbers" />
      <button type="submit">Calculate</button>
    </form>
    <div class="results">
      <p>
        The <dfn>mean</dfn> of a list of numbers is the average, calculated by
        taking the sum of all numbers and dividing that by the count of numbers.
      </p>
      <p class="bold">Mean: <span id="mean"></span></p>
      <p>
        The <dfn>median</dfn> of a list of numbers is the number that appears in
        the middle of the list, when sorted from least to greatest.
      </p>
      <p class="bold">Median: <span id="median"></span></p>
      <p>
        The <dfn>mode</dfn> of a list of numbers is the number that appears most
        often in the list.
      </p>
      <p class="bold">Mode: <span id="mode"></span></p>
      <p>
        The <dfn>range</dfn> of a list of numbers is the difference between the
        largest and smallest numbers in the list.
      </p>
      <p class="bold">Range: <span id="range"></span></p>
      <p>
        The <dfn>variance</dfn> of a list of numbers measures how far the values
        are from the mean, on average.
      </p>
      <p class="bold">Variance: <span id="variance"></span></p>
      <p>
        The <dfn>standard deviation</dfn> of a list of numbers is the square
        root of the variance.
      </p>
      <p class="bold">
        Standard Deviation: <span id="standardDeviation"></span>
      </p>
    </div>
  </body>
</html>
```

```css
body {
  margin: 0;
  background-color: rgb(27, 27, 50);
  text-align: center;
  color: #fff;
}

button {
  cursor: pointer;
  background-color: rgb(59, 59, 79);
  border: 3px solid white;
  color: white;
}

input {
  background-color: rgb(10, 10, 35);
  color: white;
  border: 1px solid rgb(59, 59, 79);
}

.bold {
  font-weight: bold;
}
```

```js
const getMean = (array) => array.reduce((acc, el) => acc + el, 0) / array.length;

--fcc-editable-region--
const testArr1 = [1, 2, 3, 4, 5];
const testArr2 = [1, 2, 3, 4, 5, 6];
const isEven = testArr2.length % 2 === 0;
console.log(isEven);


--fcc-editable-region--
const getMedian = (array) => {
  const sorted = array.sort((a, b) => a - b);
}

const calculate = () => {
  const value = document.querySelector("#numbers").value;
  const array = value.split(/,\s*/g);
  const numbers = array.map(el => Number(el)).filter(el => !isNaN(el));
  
  const mean = getMean(numbers);

  document.querySelector("#mean").textContent = mean;
}
```
