---
id: 66323433f931ca32305a11f5
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

現在，CamperBot 正嘗試創建一個函式（程式），該函式將從 `darkColorsArr` 傳回一個隨機索引。但他們遇到了以下錯誤訊息：

```js
Uncaught ReferenceError: math is not defined
```

當取用不存在的變數時，會發出 `ReferenceError`。在這種情況下，看起來 CamperBot 正嘗試使用 `math`，但 JavaScript 並沒有 `math` 物件。

修正 CamperBot 在 `math.random()` 行的錯誤，然後再次開啟主控台。

# --hints--

你應該修正 `math.random()` 行中的大小寫錯誤。 

```js
assert.match(getRandomIndex.toString(), /console\.log\(\s*darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
function getRandomIndex() {
--fcc-editable-region--
  console.log(darkColorsArr.length * math.random())
--fcc-editable-region--
}
getRandomIndex();
```
