---
id: 663250b42513ef5975599c49
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

現在，`ReferenceError` 已解決，控制檯將顯示 `0` 到 `9` 之間隨機數的正確結果。 但是，CamperBot 沒想到會看到這樣的十進制數字：

```js
0.015882899879771095
2.114596286197641
6.040964780197666
```

更新 `console` 語句，打印介於 `0` 和 `9` 之間的整數。

記得你在角色扮演遊戲中使用過一種方法，將一個數字向下舍入到最接近的整數。

# --hints--

你應該將 `darkColorsArr.length * Math.random()` 四捨五入爲最接近的整數。

```js
assert.match(getRandomIndex.toString(), /console\.log\(\s*Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
function getRandomIndex() {
  console.log(darkColorsArr.length * Math.random())
}
getRandomIndex();
--fcc-editable-region--
```
