---
id: 6650c9a94d6e13d14a043a69
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

CamperBot 正在嘗試建置一個隨機背景顏色變換器。但他們持續遇到問題，需要你幫忙偵錯 `code`。 

CamperBot 已經為專案添加了 HTML 和 CSS。但他們不明白為什麼頁面上沒有顯示任何樣式和內容。

當他們打開主控台時，他們會看到這則訊息：

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

當 JavaScript 引擎遇到無法解讀的內容時，會發生語法錯誤。在這種情況下，看起來 CamperBot 在 `darkColorsArr` 陣列中有語法錯誤。

修正 `darkColorsArr` 陣列中的語法錯誤，你就會看到內容和樣式顯示在頁面上。

# --hints--

你應該解決 `darkColorsArr` 陣列中的語法錯誤。請記得陣列成員應該以逗號分隔。

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
