---
id: 66c64210413532ee9d3bd342
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Para a próxima parte do workshop, você praticará adicionar alimentos à lista de compras.

Comece utilizando `console.log` para registrar a mensagem `"It will be nice to have some fruit to eat."`. 

# --hints--

Você deve registrar a mensagem `"It will be nice to have some fruit to eat."` no console.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
