---
id: 66c63bfa4685e4e3843fa0dc
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você continuará a aprender sobre arrays criando uma lista de compras.

Comece adicionando um `console.log` que registre a string `"Grocery shopping list"` no console.

# --hints--

Você deve ter um `console.log` no código.

```js
assert.match(code, /console\.log(.*)/);
```

O `console.log` deve produzir a mensagem `"Grocery shopping list"`. Verifique novamente se há erros de grafia e capitalização.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
