---
id: 66c6491fe4c8e0f16845425f
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora é hora de adicionar outra fruta à lista.

Usando o mesmo método de array de antes, adicione a string `"Grapes"` ao final do array `shoppingList`.

Então, adicione um `console.log` e chame a função `getShoppingListMsg` dentro do `console.log` para ver a lista atualizada registrada no console.

# --hints--

O array `shoppingList` deve agora conter dois itens.

```js
assert.lengthOf(shoppingList, 2);
```

O array `shoppingList` deve conter a string `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Você deve chamar a função `getShoppingListMsg(shoppingList)` dentro do `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
