---
id: 66c726c34ecf1e238aa9d7d5
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Nas lições anteriores, você aprendeu como adicionar elementos ao início de um array usando o método `unshift()`. 

Aqui está um lembrete de como usar o método `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Use o método `unshift()` para adicionar a string `"Vegetable Oil"` ao início do array `shoppingList`.

# --hints--

Você deve ter três itens no array `shoppingList`.

```js
assert.lengthOf(shoppingList, 3);
```

Você deve usar o método `unshift()` para adicionar `"Vegetable Oil"` ao início do array `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
