---
id: 694648acde178bb8202d9517
title: Passaggio 10
challengeType: 20
dashedName: step-10
---

# --description--

L'output è `True`, confermando che `score` è un `float`.

Completa il pagellino stampando il valore `score` insieme al suo tipo di dato usando una singola istruzione `print()`.

# --hints--

Dovresti stampare `score` e `type(score)` usando una virgola come separatore.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
