---
id: 694648acde178bb8202d9515
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

Ora, aggiungi il punteggio dello studente.

Dichiara una variabile chiamata `score` e assegnale il valore `80.5`.

Anche se sia `age` che `score` sono numeri, potrebbero non essere dello stesso tipo. Python fornisce una funzione chiamata `isinstance()` per verificarlo.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Usa `isinstance()` per verificare se `score` è un `int` e stampa il risultato nel terminale come mostrato nell'esempio sopra.

# --hints--

Dovresti avere una variabile `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

La variabile `score` dovrebbe archiviare il valore `80.5`. Non racchiudere il valore tra virgolette.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Dovresti stampare il risultato di `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
