---
id: 694648acde178bb8202d9516
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

L'output è `False`, che mostra che `score` non è un `int`.

Un altro tipo comune di numero in Python è `float`, che rappresenta un numero con decimali. Sostituisci `int` con `float` nella chiamata `isinstance()` esistente per confermarlo.

# --hints--

Dovresti sostituire `int` con `float` nella chiamata `isinstance(score, int)` esistente.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
