---
id: 69461e4f98594012939bc1c6
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo workshop, eserciterai i tipi di dati Python costruendo una semplice stampante per pagelle.

Come hai imparato nelle lezioni precedenti, le variabili si assegnano in questo modo:

```py
greeting = 'Hello World'
```

Nell'esempio sopra, la variabile `greeting` ha il valore di una stringa, che è una sequenza di caratteri racchiusa tra virgolette.

Crea una variabile chiamata `name` e assegnale la stringa `Alice`. Ricorda di racchiudere il valore tra virgolette singole o doppie come mostrato nell'esempio.

**Importante:** Python usa l'indentazione (gli spazi all'inizio di una riga) per organizzare il codice. Per questo workshop, assicurati che non ci siano spazi extra all'inizio di ogni riga di codice. Aggiungere spazi extra causerà un `IndentationError` e impedirà l'esecuzione del tuo codice.

# --hints--

Dovresti avere una variabile `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

La variabile `name` dovrebbe archiviare il valore `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
