---
id: 694648acde178bb8202d9514
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

Ora devi aggiungere l'età dello studente alla pagella. Per questo userai un intero, uno dei tipi di dati numerici in Python.

Dichiara una variabile chiamata `age` e assegnale il valore intero `20`.

Poi, stampa insieme il valore e il tipo di dato di `age` separati da una virgola. Controlla l'output nel terminale che mostra il valore di `age` e il suo tipo come `int` (intero).

# --hints--

Dovresti avere una variabile `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

La variabile `age` dovrebbe archiviare il valore `20`. Non racchiudere il valore tra virgolette.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Dovresti stampare `age` e `type(age)` usando una virgola come separatore.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
