---
id: 694648abde178bb8202d9511
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Ora dovresti vedere il nome dello studente stampato nel terminale.

Python fornisce una funzione chiamata `type()` che puoi usare per verificare il tipo di un valore.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

Nell'esempio sopra, l'output `<class 'str'>` significa che la variabile passata alla funzione `type()` è una stringa.

Usa la funzione `type()` con `name` come argomento e stampa l'output come nell'esempio. Controlla l'output nel terminale che mostra che `name` è di tipo `str` (stringa).

# --hints--

Dovresti stampare il risultato della chiamata a `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
