---
id: 694648acde178bb8202d9512
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Il pagellino dovrebbe mostrare anche se lo studente è attualmente iscritto. Questo può essere rappresentato usando un valore booleano.

I valori booleani rappresentano una condizione sì o no e sono spesso usati per prendere decisioni nel codice. Esistono solo due valori booleani: `True` e `False`.

Dichiara una variabile chiamata `is_student` e assegnale il valore `True`.

# --hints--

Dovresti avere una variabile `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

La variabile `is_student` dovrebbe archiviare il valore `True`. Non racchiudere il valore tra virgolette.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
