---
id: 694648acde178bb8202d9513
title: Passaggio 6
challengeType: 20
dashedName: step-6
---

# --description--

Il nome dello studente dovrebbe seguire lo stesso formato degli altri dettagli.

Rimuovi gli output precedenti della variabile `name`. Poi, stampa `name` e `type(name)` insieme su una riga separati da una virgola come nel passaggio precedente.

# --hints--

Non dovresti avere `print(name)` nel tuo codice.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Non dovresti avere `print(type(name))` nel tuo codice.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Dovresti stampare `name` e `type(name)` nella stessa chiamata `print()` usando una virgola come separatore.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
