---
id: 6947a3e516bb8a4a681274a7
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

La funzione `print()` può mostrare più valori contemporaneamente. Separa i valori con una virgola (`,`) per stamparli sulla stessa riga.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Stampa sia `is_student` che `type(is_student)` sulla stessa riga usando una virgola `,` come mostrato nell'esempio. Poi, controlla l'output nel terminale che mostra il valore di `is_student` e il suo tipo come `bool` (booleano).

# --hints--

Dovresti stampare `is_student` e `type(is_student)` usando una virgola come separatore.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
