---
id: 6947a0c2b6de035dd8d3eea1
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Puoi stampare il valore di una variabile usando la funzione `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Imparerai di più sulle funzioni nelle prossime lezioni. Per ora, sappi che una funzione è un blocco di codice riutilizzabile che può essere chiamato, o invocato, per eseguire il suo codice, e a cui possono essere passati degli argomenti.

Nell'esempio sopra, `print(greeting)` è una chiamata di funzione, e `greeting` è l'argomento della funzione.

Fai riferimento all'esempio e stampa la variabile `name`. Controlla l'output nel terminale.

# --hints--

Dovresti chiamare `print()` con `name` come suo argomento.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
