---
id: a202eed8fc186c8434cb6d61
title: Tengeneza Kigeuzi cha Mfuatano wa Herufi
challengeType: 26
dashedName: build-a-string-inverter
---

# --description--

Katika maabara hii, utatengeneza kigeuzi rahisi cha mfuatano wa herufi kinachogeuza mfuatano wa herufi uliotolewa.

Kwa mfano, `"hello"` inapaswa kuwa `"olleh"`.

**Lengo:** Timiza hadithi za mtumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za mtumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `reverseString` kinachopokea mfuatano wa herufi kama hoja.
2. Kitendakazi kinapaswa kurudisha mfuatano wa herufi uliogeuzwa nyuma.

# --hints--

Unapaswa kuwa na kitendakazi kinachoitwa `reverseString`.

```js
assert.isFunction(reverseString);
```

`reverseString` inapaswa kupokea mfuatano wa herufi kama hoja.

```js
assert.match(reverseString.toString(), /\s*function\s+reverseString\s*\(\s*\w+\s*\)/);
```

`reverseString("hello")` inapaswa kurudisha mfuatano wa herufi.

```js
assert.isString(reverseString('hello'));
```

`reverseString("hello")` inapaswa kurudisha mfuatano wa herufi `olleh`.

```js
assert.strictEqual(reverseString('hello'), 'olleh');
```

`reverseString("Howdy")` inapaswa kurudisha mfuatano wa herufi `ydwoH`.

```js
assert.strictEqual(reverseString('Howdy'), 'ydwoH');
```

`reverseString("Greetings from Earth")` inapaswa kurudisha mfuatano wa herufi `htraE morf sgniteerG`.

```js
assert.strictEqual(
  reverseString('Greetings from Earth'),
  'htraE morf sgniteerG'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reverseString(str) {
  return str.split('').reverse().join('');
}

reverseString('hello');
```
