---
id: 68e6698b29e562d21b1d260c
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utaunda orodha ya vivinjari vikuu vya mtandao. Muundo wa msingi wa HTML umeandaliwa kwa ajili yako.

Anza kwa kuongeza kichwa cha sehemu kwenye ukurasa wako kinachosoma `List of Major Web Browsers` kwa kutumia kipengele cha `h1` ndani ya kipengele cha `body`.

# --hints--

Unapaswa kuwa na kipengele cha `h1` ndani ya kipengele cha `body`.

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi `List of Major Web Browsers`.

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>List of Browsers and Descriptions</title>
  </head>
  <body>
  --fcc-editable-region--
    
  --fcc-editable-region--
  </body>
</html>
```
