---
id: 68ecfa1003d88990bb96a33e
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Sasa ni wakati wa kuongeza kivinjari chako cha pili kwenye orodha.

Chini ya kipengele chako cha `Google Chrome`, ongeza kipengele kingine cha `dt` chenye maandishi ya `Firefox`.

Chini ya kipengele chako cha `dt`, ongeza kipengele cha `dd` chenye maandishi yafuatayo:

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

Unapaswa kuwa na vipengele viwili vya `dt` katika msimbo wako.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

Kipengele chako cha pili cha `dt` kinapaswa kuwa na maandishi ya `Firefox`.

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

Unapaswa kuwa na vipengele viwili vya `dd` katika msimbo wako.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

Kipengele chako cha pili cha `dd` kinapaswa kuwa na maandishi ya `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```
