---
id: 68ecfa1003d88990bb96a33d
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Orodha za maelezo zinaundwa na orodha ya maneno na maelezo. Kipengele cha `dt` kinawakilisha neno la maelezo, wakati kipengele cha `dd` kinawakilisha maelezo ya maelezo.

Kivinjari cha kwanza utakachoongeza kwenye orodha yako ya maelezo kitakuwa Google Chrome.

Ndani ya kipengele chako cha `dl`, ongeza kipengele cha `dt` chenye maandishi ya `Google Chrome`.

Chini ya kipengele chako cha `dt`, ongeza kipengele cha `dd` chenye maandishi yafuatayo:

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

Unapaswa kuwa na kipengele cha `dt` ndani ya kipengele chako cha `dl`.

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

Kipengele chako cha `dt` kinapaswa kuwa na maandishi ya `Google Chrome`.

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

Unapaswa kuwa na kipengele cha `dd` chini ya kipengele chako cha `dt`.

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

Kipengele chako cha `dd` kinapaswa kuwa na maandishi ya `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```
