---
id: 68ecfa0f03d88990bb96a33c
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Kama unavyokumbuka kutoka kwenye funzo la awali, orodha za maelezo hutumika kuwasilisha maneno na ufafanuzi kwa mpangilio rahisi na rahisi kusoma.

Hapa kuna mfano:

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

Chini ya kipengele cha `h1`, tengeneza kipengele cha `dl`. Hiki kitashikilia orodha ya vivinjari.

# --hints--

Unapaswa kuwa na kipengele cha `dl`.

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

Kipengele chako cha `dl` kinapaswa kuwa chini ya kipengele cha `h1`.

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    --fcc-editable-region--
    
    --fcc-editable-region--
  </body> 
</html>
```

