---
id: 69038dd51a41ae215fcc030a
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Kivinjari cha tatu utakachoongeza kwenye orodha kitakuwa kivinjari cha mtandao cha Safari.

Ongeza kipengele kingine cha `dt` chenye maandishi ya `Safari`.

Chini ya kipengele chako cha `dt`, ongeza kipengele cha `dd` chenye maandishi yafuatayo:

```md
This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.
```

# --hints--

Unapaswa kuwa na vipengele vitatu vya `dt` katika msimbo wako.

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 3);
```

Kipengele chako cha tatu cha `dt` kinapaswa kuwa na maandishi ya `Safari`.

```js
const dt = document.querySelectorAll('dt')[2];
assert.strictEqual(dt?.textContent.trim(), "Safari");
```

Unapaswa kuwa na vipengele vitatu vya `dd` katika msimbo wako.

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 3);
```

Kipengele chako cha tatu cha `dd` kinapaswa kuwa na maandishi ya `This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.`

```js
const ddElement = document.querySelectorAll('dd')[2];
assert.strictEqual(ddElement?.textContent.trim(), "This browser was developed by Apple and is the default browser for iPhone, iPad and Mac devices.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
  <head> 
    <meta charset="UTF-8"> 
    <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
    <title>List of Browsers and Descriptions</title> 
  </head> 
  <body> 
    <h1>List of Major Web Browsers</h1> 
    <dl>
      <dt>Google Chrome</dt>
      <dd>This is a free web browser developed by Google and first released in 2008.</dd>

      <dt>Firefox</dt>
      <dd>This is a free web browser developed by the Mozilla Corporation and first created in 2004.</dd>

        --fcc-editable-region--
      
        --fcc-editable-region--
    </dl>
  </body> 
</html>
```
