---
id: 612e83ec2eca1e370f830511
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`head` 要素内に `link` 要素を追加してください。 その `rel` 属性を `stylesheet` に、`href` 属性を `styles.css` に設定してください。

# --hints--

コードには `link` 要素が 1 つ必要です。

```js
assert.isNotNull(document.querySelector('link'));
```

既存の `head` タグを変更しないようにしてください。 終了タグを削除していないか確認してください。

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

`link` 要素が 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

`link` 要素は `head` 要素内に置く必要があります。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

`link` 要素は値が `stylesheet` に設定されている `rel` 属性を持つ必要があります。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素は値が `styles.css` に設定されている `href` 属性を持つ必要があります。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    --fcc-editable-region--
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    --fcc-editable-region--
  </head>
  <body>
    <div id="piano">
      <div class="keys">
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>

        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
        <div class="key black--key"></div>
      </div>
    </div>
  </body>
</html>
```

```css

```
