---
id: 612e78af05201622d4bab8aa
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

Time to start working on the piano. Create a `div` element within your `body` element with the `id` set to `piano`.

# --hints--

新しい `div` 要素を作成してください。

```js
const div = document.querySelector('div');
assert.exists(div);
```

`div` 要素は `body` 要素の中にある必要があります。

```js
const div = document.querySelector('div');
assert(div?.parentElement?.localName === 'body');
```

`div` 要素の `id` の値は `piano` に設定される必要があります。

```js
const div = document.querySelector('div');
assert(div?.getAttribute('id') === 'piano');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Piano</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
  <body>
  --fcc-editable-region--

  --fcc-editable-region--
  </body>
</html>
```

```css

```
