---
id: 6852db755c9c1274d75aae9a
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、詩の各行に隠されたピンの数字を抽出するピン抽出器を作成します。

まず、`pin_extractor` という名前の関数を、`poem` というパラメータ付きで作成してください。空の関数はエラーになるので、一時的に関数内に `pass` を追加して有効にしてください。

# --hints--

`pin_extractor` という名前の関数があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 関数は `poem` という名前のパラメータを持つべきです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
