---
id: 6852db755c9c1274d75aae9a
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、PIN抽出器を作成します。PINの数字は詩の各行に隠されています。

`pin_extractor`という名前の関数を、`poem`というパラメータ付きで作成することから始めてください。関数に本体がないとエラーになるので、有効にするために関数内に一時的なプレースホルダーとして`pass`を追加してください。

# --hints--

`pin_extractor` という名前の関数があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 関数は `poem` という名前のパラメータを持つべきです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
