---
id: 6852dd8486c7ee0171322950
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

デコード中に秘密のコードを保持する変数が必要なので、`pass` を削除し、関数内に `secret_code` という変数を作成して空の文字列を代入してください。

# --hints--

`secret_code` 関数の中に `pin_extractor` という変数があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

変数 `secret_code` には空の文字列が代入されている必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

コード内に `pass` はもう存在しないはずです。

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
