---
id: 68531fbb934d735415053274
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

関数の呼び出しを追加したので、関数内で発生したエラーはターミナルに表示され、`print` を使って値を確認することもできます。

ピンの数字が各行に隠れているので、関数内で `split` メソッドを使って文字列を行のリストに分割してください。改行文字（`\n`）で分割し、結果のリストを `lines` という変数に代入してください。

# --hints--

`lines` 関数内に `pin_extractor` という名前の変数があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

変数 `lines` は `poem.split('\n')` の値を持っているはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
