---
id: 68530231c05b93324e071104
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

関数を作成している間、デバッグ用に関数を呼び出すことができると便利なので、関数の外にこの詩を含む`poem`変数を作成してください。

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

複数行の文字列を使うことができます。

そして、その`poem`変数を引数として関数を呼び出してください。

# --hints--

`poem`という名前のグローバル変数を用意してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

`poem`変数の値は`Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`と等しくなるはずです。

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

`pin_extractor`関数を`poem`を引数として呼び出してください。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
