---
id: 685535bf527b2c77236b3a9a
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

ピンのn番目の数字は、n番目の行のn番目の単語の長さとして隠されています。

n番目の単語の長さを見つけるために、次のステップは詩の行を単語のリストに分割することです。ループ内で、変数`words`を作成し、`line`を`split`メソッドを使って単語に分割した値を代入してください。

その後、`print`を引数にして`words`の呼び出しを追加してください。

# --hints--

ループ内に`words`変数があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

`words`変数は`line.split()`の値を持つはずです。

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

`print(words)`があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
