---
id: 68557bd29f7f3c1880f98836
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

ループの最後の行で、`print` の呼び出しを更新して、`words` リストのインデックス `line_index` にある単語を出力するようにしてください。

# --hints--

ループの最後に `print(words[line_index])` があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words)
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
