---
id: 685578aae65e4106ad2b4300
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

前のレッスンで学んだように、`enumerate` 関数を使うと、イテラブルを繰り返す際にインデックスを追跡できます。

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

秘密のコードのn番目の数字は、n番目の行のn番目の単語から来ているので、各行の番号を知る必要があります。

ループを`enumerate(lines)`に繰り返すように変更し、`line`の前に`line_index`という別のループ変数を追加してください。

また、`print`の呼び出しを`print(line_index, line)`に更新してください。

# --hints--

ループは`enumerate(lines)`を繰り返す必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

ループ変数として`line_index, line`を使う必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

ループ内に`print(line_index, line)`を含めてください。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
