---
id: 68553423d775cc60134ac8bb
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

各行を独立して処理する必要があります：`lines` を繰り返すループを作成し、`line` をループ変数として使ってください。ループの中で `line` を出力してください。

# --hints--

`for` 関数内に `pin_extractor` ループがあるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

ループは `lines` を繰り返す必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

ループ変数は `line` であるべきです（`for line in lines`）。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

ループの中に `print(line)` があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
