---
id: 68559f3561d7df287586142e
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

ピン抽出機能は動作していますが、もっと効率的にできます！複数の詩から同時にピンを抽出できるようにしましょう！

グローバルスコープにさらに2つの詩、`poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'` と `poem3 = 'There\nonce\nwas\na\ndragon'` を追加し、今は `print` の `pin_extractor` 呼び出しをコメントアウトしてください。


# --hints--

変数 `poem2` と `poem3` は存在し、正しい値を持っている必要があります。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

`print(pin_extractor(poem))` の呼び出しはコメントアウトしてください。

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
