---
id: 6855918e87c33c6979d21d3f
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

関数は詩からコードを抽出する作業をうまくこなしていますが、考慮すべき例外ケースがあります。詩の行が予想より短い場合、関数がエラーになる可能性があります（詩から単語`bright`を削除して試すことができます）。

詩の行`secret_code += str(len(words[line_index]))`を、`if`リストに十分な単語があるかどうかを確認する`words`文に入れてください。

# --hints--

ループの中に`if`文を入れる必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

if文の条件は`len(words) > line_index`にしてください。


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

`if`文の中に`secret_code += str(len(words[line_index]))`を入れてください。


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
