---
id: 6859a037b3d6721c7fbb8055
title: ステップ 18
challengeType: 20
dashedName: step-18
---

# --description--

ループの前に、空のリストを初期値とする新しい変数`secret_codes`を作成してください。

次に、`return secret_code`の行を、`secret_code`を`secret_codes`リストに追加する行に置き換えてください。

# --hints--

`secret_codes`変数は`pin_extractor`関数の中にあるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

変数`secret_codes`は空のリストの値を持つべきです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

ループの最後には`secret_codes.append(secret_code)`があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

`secret_code`を返してはいけません。

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```


# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
