---
id: 68592730aaf34b61ae279168
title: ステップ 10
challengeType: 20
dashedName: step-10
---

# --description--

`len` から返される長さは整数なので、`str` 関数を使って文字列に変換し、`secret_code` 文字列に追加できるようにしてください。

`print` の呼び出しを更新して、長さを文字列として出力してください。

# --hints--

ループの最後に `print(str(len(words[line_index])))` があるはずです。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
