---
id: 6859290fc8e07e6b2a8d8523
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

2つの`print`の呼び出しを削除し、関数の最後の行に`return`を返す`secret_code`文を追加してください。

# --hints--

`print`の呼び出しはあってはいけません。

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

`pin_extractor`関数は最後に`secret_code`を返す必要があります。

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)`は`5262`を返すべきです。

```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
--fcc-editable-region--
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))
    
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
