---
id: 6564aee9c077774ea49c3faf
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

Der Merge-Sort-Algorithmus führt drei Aktionen aus:

- Teilen einer unsortierten Folge von Elementen in Unterteile
- Sortieren der Elemente in den Unterteilen
- Zusammenführen der sortierten Unterteile

Das oben Genannte geschieht rekursiv, bis die Unterteile zu der vollständig sortierten Folge zusammengeführt sind. Beginnen Sie mit dem Teilen der Folge.

Ersetzen Sie zuerst das Schlüsselwort `pass` durch eine Variable `middle_point`. Verwenden Sie den ganzzahligen Divisionsoperator (`//`), um die Länge der `array`-Liste durch 2 zu teilen, und weisen Sie das Ergebnis Ihrer neuen `middle_point`-Variable zu. Denken Sie daran, Ihren Code einzurücken.

# --hints--

Sie sollten `pass` nicht im Funktionskörper haben.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

Sie sollten eine `middle_point`-Variable deklarieren und `len(array) // 2` zuweisen.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    pass
--fcc-editable-region--
```
