---
id: 656639771fed09b5c0e8fe71
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Nachdem Sie die `array`-Liste in zwei separate Listen aufgeteilt haben, teilen Sie jede Liste weiter auf, bis jedes Element in seiner eigenen Liste steht. Eine Liste mit nur einer Zahl ist immer sortiert.

Rufen Sie dazu rekursiv `merge_sort` innerhalb Ihrer Funktion auf und übergeben Sie `left_part` als Argument des Aufrufs.

# --hints--

Sie sollten `merge_sort` mit dem Argument `left_part` am Ende Ihres Funktionskörpers aufrufen.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
--fcc-editable-region--
    
--fcc-editable-region--
```
