---
id: 656680b0fc79f2c38a34d90e
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Jetzt ist es Zeit, die Listen (`left_part` und `right_part`) zu sortieren und in die ursprüngliche Liste zusammenzuführen.

Sie können dies tun, indem Sie die Elemente beider Listen vergleichen und das kleinere Element in die Hauptliste übernehmen. Diesen Vergleich führen Sie für alle Indizes in `left_part` und `right_part` durch.

Erstellen Sie drei Variablen: `left_array_index`, `right_array_index` und `sorted_index` und setzen Sie deren Werte auf `0`. Diese Variablen helfen Ihnen, während des Sortierprozesses jeden Index im Blick zu behalten.

# --hints--

Sie sollten eine Variable namens `left_array_index` innerhalb Ihrer `merge_sort` Funktion haben

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

Der Wert von `left_array_index` sollte `0` sein.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

Sie sollten eine Variable namens `right_array_index` innerhalb Ihrer `merge_sort` Funktion haben

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

Der Wert von `right_array_index` sollte `0` sein.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

Sie sollten eine Variable namens `sorted_index` innerhalb Ihrer `merge_sort` Funktion haben

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

Der Wert von `sorted_index` sollte `0` sein.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
--fcc-editable-region--
    
--fcc-editable-region--
```
