---
id: 656627b47bd2d2a4afbc945d
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Im vorherigen Schritt haben Sie den Mittelpunkt ermittelt. Sie können ihn verwenden, um `array` in zwei Teile zu teilen und jeden Teil neuen Variablen zuzuweisen.

Verwenden Sie die Slice-Syntax, um die linke Hälfte von `array` zu extrahieren und sie einer Variablen namens `left_part` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `left_part` in Ihrer `merge_sort`-Funktion haben.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

Sie sollten die Slice-Syntax verwenden, um die linke Hälfte des Arrays zu extrahieren und sie der Variablen `left_part` zuzuweisen.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
--fcc-editable-region--
    
--fcc-editable-region--
```
