---
id: 656710d1e0ec62253426db24
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Die von Ihnen erstellte `while`-Schleife vergleicht ein Element aus `left_part` mit einem anderen in `right_part` und fügt dann das kleinere Element zur Hauptliste `array` hinzu.

Dieser Vorgang wird so lange fortgesetzt, bis keine Elemente mehr zum Vergleichen übrig sind. Aber `left_part` kann noch Elemente enthalten, während `right_part` keine mehr hat, und umgekehrt.

Erstellen Sie eine weitere `while`-Schleife, die ausgeführt wird, wenn `left_array_index` kleiner als `len(left_part)` ist. In den nächsten Schritten verwenden Sie diese, um die verbleibenden Elemente in `left_part` in die Liste `array` zu kopieren.

# --hints--

Ihre `while`-Schleife sollte diese Bedingung haben: `left_array_index < len(left_part)`. Vergessen Sie nicht, das Schlüsselwort `pass` im Schleifenkörper zu verwenden.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

Sie sollten das Schlüsselwort `pass` im Körper Ihrer `while`-Schleife haben.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
