---
id: 656706afd65547d22bee0b68
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Ersetzen Sie innerhalb der `while`-Schleife `pass` durch eine `if`-Anweisung, die überprüft, ob das Element in `left_part` kleiner ist als das Element in `right_part`.

Verwenden Sie das Schlüsselwort `pass` im Körper der `if`-Anweisung.

# --hints--

Sie sollten das Schlüsselwort `pass` durch eine `if`-Anweisung mit dieser Bedingung ersetzen: `left_part[left_array_index] < right_part[right_array_index]`. Fügen Sie das Schlüsselwort `pass` in den Körper der `if`-Anweisung ein.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
