---
id: 656709e50ed928da35260276
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

In einem vorherigen Schritt haben Sie das Element in der `left_part` der `array`-Liste zugewiesen, weil es kleiner war. Aber das wird nicht immer der Fall sein. In einigen Vergleichsfällen könnte das Element rechts kleiner sein.

Erstellen Sie eine `else`-Klausel, die ausgeführt wird, wenn das Element in `left_part` nicht kleiner ist als das Element in `right_part`.

Weisen Sie innerhalb des `else`-Blocks `right_part[right_array_index]` `array[sorted_index]` zu.

# --hints--

Sie sollten `array[sorted_index] = right_part[right_array_index]` im Körper Ihrer `else`-Klausel haben.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
--fcc-editable-region--
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
```
