---
id: 656702f8b4cbd8cbf0a433c6
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Als Nächstes müssen Sie eine Schleife erstellen, die so lange fortgesetzt wird, wie in beiden `left_part` und `right_part` noch Elemente vorhanden sind.

Erstellen Sie dafür eine `while`-Schleife mit zwei Bedingungen: eine, die überprüft, ob der `left_array_index` kleiner als die Länge von `left_part` ist, und eine weitere, die überprüft, ob `right_array_index` kleiner als die Länge von `right_part` ist.

# --hints--

Ihre `while`-Schleife sollte prüfen, ob `left_array_index` kleiner als die Länge von `left_part` ist und ob `right_array_index` kleiner als die Länge von `right_part` ist.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

Verwenden Sie das Schlüsselwort `pass` im Schleifenkörper.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
