---
id: 656707e11e671ed54c96f7ec
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Wenn die `if`-Bedingung zu `True` ausgewertet wird, bedeutet das, dass das Element in der `left_part`-Liste kleiner ist als das Element, mit dem es in der `right_part`-Liste verglichen wird.

In diesem Fall können Sie das an `left_array_index` gefundene Element an `left_part` innerhalb des sortierten Arrays einfügen.

Innerhalb des `if`-Blocks entfernen Sie `pass` und weisen `left_part[left_array_index]` `array[sorted_index]` zu.

# --hints--

Sie sollten `pass` im Körper der `array[sorted_index] = left_part[left_array_index]`-Anweisung durch `if` ersetzen.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_body().is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            pass

--fcc-editable-region--
```
