---
id: 6569b743630ee592a65a7e2f
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Bevor Sie die `merge_sort()`-Funktion testen, müssen Sie einen Basisfall erstellen, der die Ausführung der Funktion stoppt, wenn die Länge von `array` kleiner oder gleich `1` ist.

Dieser Basisfall stoppt den Rekursionsaufruf. Ohne ihn würde der Merge-Sort-Vorgang weiterlaufen, selbst wenn die Liste bereits sortiert ist oder keine Elemente enthält.

Direkt nach der Funktionsdeklaration erstellen Sie eine `if`-Anweisung mit dieser Bedingung: `len(array) <= 1`. Innerhalb des Körpers der `if`-Anweisung fügen Sie eine `return`-Anweisung hinzu, um die Ausführung der Funktion zu stoppen.

# --hints--

Sie sollten eine `if`-Anweisung erstellen, um zu prüfen, ob `len(array) <= 1`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

Sie sollten eine `return`-Anweisung im Körper der `if`-Anweisung haben.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_ifs()[0].find_body().has_stmt("return")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    
--fcc-editable-region--
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
