---
id: 6569b20f829b7e69d43c232a
title: Schritt 18
challengeType: 20
dashedName: step-18
---

# --description--

Das Letzte, was für die `while`-Schleife zu tun ist, ist, zum nächsten Index im sortierten Array zu wechseln.

Fügen Sie `1` zum Wert von `sorted_index` hinzu.

# --hints--

Sie sollten den aktuellen Wert von `sorted_index` innerhalb der `while`-Schleife um eins erhöhen. Beachten Sie, dass die Erhöhung nach der Zuweisung des linken Teils zum sortierten Array erfolgen sollte.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
