---
id: 6569b0a63423f65dd30888df
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Entfernen Sie das Schlüsselwort `pass`. Weisen Sie im Codeblock der `while`-Schleife `left_part[left_array_index]` `array[sorted_index]` zu.

# --hints--

Sie sollten `array[sorted_index] = left_part[left_array_index]` im Schleifenkörper haben. Vergessen Sie nicht, das Schlüsselwort `pass` im Schleifenkörper zu entfernen.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].find_body()[0].is_equivalent("array[sorted_index] = left_part[left_array_index]")`))
})
```

Sie sollten `pass` nicht in der `while`-Schleife haben.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
--fcc-editable-region--
        pass
--fcc-editable-region--
```
