---
id: 6569b4e0bd29d17d4b53b16c
title: Schritt 19
challengeType: 20
dashedName: step-19
---

# --description--

Jetzt werden Sie die gleiche `while`-Schleifenlogik für `right_part` nachbilden.

Erstellen Sie eine `while`-Schleife, die läuft, wenn `right_array_index` kleiner als `len(right_part)` ist. Innerhalb der `while`-Schleife weisen Sie `right_part[right_array_index]` `array[sorted_index]` zu. Am Ende erhöhen Sie `right_array_index` und `sorted_index` um `1`.

# --hints--

Ihre `while`-Schleife sollte prüfen, ob `right_array_index` kleiner als `len(right_part)` ist.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

Sie sollten `right_part[right_array_index]` im Schleifenkörper `array[sorted_index]` zuweisen.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

Sie sollten `right_array_index` innerhalb der `1`-Schleife um `while` erhöhen. Beachten Sie, dass die Erhöhung nach der Zuweisung des rechten Teilelements zum sortierten Array erfolgen sollte.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

Sie sollten den aktuellen Wert von `sorted_index` innerhalb der `1`-Schleife um `while` erhöhen. Beachten Sie, dass die Erhöhung nach der Zuweisung des rechten Teilelements zum sortierten Array erfolgen sollte.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
