---
id: 6577254891048c8f2c19e961
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Rufen Sie am Ende Ihres Funktionskörpers die Funktion `merge_sort()` erneut auf. Dieses Mal übergeben Sie `right_part` als Argument beim Funktionsaufruf.

# --hints--

Sie sollten `merge_sort()` mit dem Argument `right_part` am Ende Ihres Funktionskörpers aufrufen.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
