---
id: 657f3a661730891aa64f3568
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Verwenden Sie die Slice-Syntax, um die rechte Hälfte von `array` zu extrahieren und einer Variablen mit dem Namen `right_part` zuzuweisen.

# --hints--

Sie sollten eine Variable mit dem Namen `right_part` haben

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

Sie sollten die Slice-Syntax verwenden, um die rechte Hälfte des Arrays zu extrahieren und der Variablen `right_part` zuzuweisen.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
    
--fcc-editable-region--
```
