---
id: 657f59751d5a5c9b069ae0f3
title: Schritt 21
challengeType: 20
dashedName: step-21
---

# --description--

Wie Sie in einer vorherigen Lektion gelernt haben, können Sie die Variable `__name__` verwenden, um festzustellen, ob ein Python-Skript als Hauptprogramm ausgeführt wird oder ob es als Modul importiert wird (Code, der in einer anderen Python-Datei geschrieben ist).

Wenn der Wert von `__name__` auf `'__main__'` gesetzt ist, bedeutet das, dass das aktuelle Skript das Hauptprogramm ist und kein Modul.

In diesem Projekt verwenden Sie das aktuelle Skript als Hauptprogramm.

Erstellen Sie eine `if`-Anweisung, die überprüft, ob der Wert von `__name__` `'__main__'` ist.
Erstellen Sie innerhalb der `if`-Anweisung eine Liste namens `numbers` und weisen Sie ihr diese Werte zu: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

Ihre `if`-Anweisung sollte prüfen, ob `__name__` den Wert `'__main__'` hat.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

Sie sollten eine `numbers`-Liste mit diesen Werten haben: `[4, 10, 6, 14, 2, 1, 8, 5]`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
