---
id: 61537c9eecea6a335db6da79
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

在你的 `.header` 元素下方，創建一個新的 `div` 元素並指定它的 `class` 為 `gallery`。這個 `div` 將作為畫廊影像的容器。

在該 `.gallery` 元素內，創建九個 `img` 元素。

# --hints--

你應該在你的 `body` 元素中創建一個 `div` 元素。

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

你的新 `div` 元素應該有一個 `class`，其值設定為 `gallery`。

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

你的新 `div` 元素應該放在你的 `.header` 元素之後。

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

你的 `.gallery` 元素應該有九個 `img` 元素。

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
