---
id: 61537a8054753e2f1f2a1574
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `head` 元素中，為你的 `meta` 頁籤添加 `name` 設為 `viewport` 且 `content` 設為 `width=device-width, initial-scale=1`。

同時添加一個 `meta` 頁籤，並將 `charset` 設定為 `UTF-8`。

# --hints--

你應該有兩個 `meta` 元素。

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

一個 `meta` 元素應該設定 `name` 為 `viewport`，並將 `content` 設定為 `width=device-width, initial-scale=1.0`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content').match(/width=device-width,\s?initial-scale=1(.0)?/) && !m?.getAttribute('charset'));
assert.exists(target);
```

你的另一個 `meta` 元素應該將 `charset` 屬性設為 `UTF-8`。

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>

  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
