---
id: 61537bb9b1a29430ac15ad38
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `head` 元素中，為你的頁面添加一個 `title` 元素，文字設定為 `Photo Gallery`，並添加一個 `link` 元素將你的 `styles.css` 檔案加入頁面。

# --hints--

你的 `link` 元素應該有一個 `href` 屬性，其值為 `styles.css`。

```js
assert.match(code, /<link[\s\S]*?href=('|"|`)(\.\/)?styles\.css\1/)
```

你的程式碼應該有一個 `title` 元素。

```js
const title = document.querySelector('title');
assert.exists(title);
```

你的專案應該有一個 `Photo Gallery` 的標題。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'photo gallery')
```

請記住，標題的大小寫和拼寫很重要。

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Photo Gallery');
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    
  </head>
  <body>
  </body>
</html>
--fcc-editable-region--
```

```css

```
