---
id: 61537d86bdc3dd343688fceb
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

接著，根據每個 `img` 在文件中的順序，為其添加 `src` 屬性。第一個 `img` 的 `src` 應該是 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。其餘的應該相同，但將 `1` 替換成該 `img` 在文件中的編號。

# --hints--

你的所有九個 `img` 元素都應該有一個 `src` 屬性。

```js
const images = [...document.querySelectorAll('img')];
assert(images.every(image => image.getAttribute('src')));
```

你的第一個 `img` 元素的 `src` 屬性值應設為 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`。

```js
assert(document.querySelectorAll('img')?.[0]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

你的第二個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[1]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

你的第三個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[2]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

你的第四個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[3]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

你的第五個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[4]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

你的第六個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[5]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

你的第七個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[6]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

你的第八個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` 設為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[7]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

你的第九個 `img` 元素應該將 `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` 設定為 `src` 屬性值。

```js
assert(document.querySelectorAll('img')?.[8]?.getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
--fcc-editable-region--
    <div class="gallery">
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
    </div>
--fcc-editable-region--
  </body>
</html>
```

```css

```
