---
id: 6153a3952facd25a83fe8083
title: 第 23 步
challengeType: 0
dashedName: step-23
---

# --description--

`::after` 偽元素會創建一個作為所選元素最後一個子元素的元素。你可以用它在最後一個影像後添加一個空元素。如果你給它與影像相同的 `width`，當畫廊處於兩行佈局時，它會將最後一個影像向左推。現在它位於中央，因為你在彈性容器上設定了 `justify-content: center`。

範例：

```CSS
.container::after {
  content: "";
  width: 860px;
}
```

使用 `.gallery` 元素上的 `::after` 偽元素創建一個新的選擇器。為 `content` 屬性添加設定為空字串 `""`，並為 `width` 屬性設定 `350px`。

# --hints--

你的 `.gallery::after` 選擇器的 `content` 屬性應該設為空字串 `""`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.content === "\"\"");
```

你的 `.gallery::after` 選擇器的 `width` 屬性應設定為 `350px`。

```js
assert(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.width === "350px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
