---
id: 61537485c4f2a624f18d7794
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

從你的標準 HTML 範本開始。為你的文件添加 `DOCTYPE` 宣告、`html` 元素、`head` 元素和 `body` 元素。

為你的開啟 `<html>` 頁籤添加 `lang` 屬性，並將值設為 `en`。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該將 `html` 定義為文件型別。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

你的 `html` 元素應該有一個開啟的頁籤。

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

你的 `html` 元素應該有一個結束標籤。

```js
assert(code.match(/<\/html\s*>/));
```

你的開啟 `<html>` 頁籤應該有 `lang` 屬性，且值為 `en`。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同層元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
