---
id: 6494da0daf5df5197963671d
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

注意藍色影像邊框如何超出紅色畫廊邊框。這是因為瀏覽器計算容器元素大小的方式所致。

`box-sizing` 屬性用來設定此行為。預設會使用 `content-box` 模型。在此模型中，當元素有特定寬度時，該寬度僅根據元素的內容計算。內距和邊框值會加到總寬度中，因此元素會擴展以容納這些值。

嘗試使用全域的 `*` 選擇器明確地將 `box-sizing` 設定為 `content-box`。此時，你不會看到任何變化，因為你使用的是預設值。

# --hints--

你應該有一個 `*` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*'));
```

你的 `*` 選擇器應該有一個 `box-sizing` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing);
```

你的 `*` 選擇器應該將其 `box-sizing` 屬性設為 `content-box`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing, 'content-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
  ```

```css
--fcc-editable-region--

--fcc-editable-region--

.gallery {
  border: 5px solid red;
  width: 50%;
}

img {
  width: 100%;
  border: 5px solid blue;
  padding: 5px;
}
```
