---
id: 64dd076a1b33c86d84b4232f
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

為了更好地視覺化你的元素大小，添加邊框會很有幫助。

為你的 `.gallery` 元素設定 `width` 為 `50%`，並將 `border` 設定為 `5px solid red`。

然後給你的 `img` 元素設定 `width` 為 `100%`、`padding` 為 `5px`，以及 `border` 為 `5px solid blue`。

# --hints--

你應該有一個 `.gallery` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

你的 `.gallery` 選擇器應該有一個 `width` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

你的 `.gallery` 選擇器應該有一個設定為 `50%` 的 `width` 屬性。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

你的 `.gallery` 選擇器應該有一個 `border` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

你的 `.gallery` 選擇器應該設定 `border` 屬性為 `5px solid red`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

你應該有一個 `img` 選擇器。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

你的 `img` 選擇器應該有一個 `width` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

你的 `img` 選擇器應該設定 `width` 屬性為 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

你的 `img` 選擇器應該有一個 `padding` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

你的 `img` 選擇器應該設定 `padding` 屬性為 `5px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

你的 `img` 選擇器應該有 `border` 屬性。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

你的 `img` 選擇器應該設定 `border` 屬性為 `5px solid blue`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
