---
id: 685976a92ffab8b2dc744b72
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

在你的 `.tabs` 元素內，添加一個 `div` 元素，其 `role` 屬性設為 `"tablist"`，並且 `aria-labelledby` 屬性設為 `"tabs-title"`。  

此處需要 `aria-labelledby` 屬性來將標籤列表與可見標籤關聯，協助像螢幕閱讀器等輔助技術向使用者宣告標籤列表的用途。

# --hints--

你應該在你的 `.tabs` 元素內有一個 `div` 元素。

```js
assert.exists(document.querySelector(".tabs div"));
```

你的 `div` 元素應該有一個 `role` 屬性設定為 `"tablist"`。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

你的 `div` 元素應該有一個 `aria-labelledby` 屬性，設定為 `"tabs-title"`。

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
