---
id: 68599d0fcd5f6ac313d05908
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

當使用者選擇一個分頁時，應將該分頁的 `aria-selected` 屬性更新為 `"true"`，並將所有其他分頁的 aria-selected 屬性更新為 `"false"`。這會顯示目前作用中的分頁，並確保輔助技術的可存取性。

預設情況下，應該選取第一個分頁。你可以透過為第一個 `button` 元素添加 `aria-selected` 屬性並將其值設為 `"true"` 來達成此目的。

# --hints--

你的第一個 `button` 元素應該有一個 `aria-selected` 屬性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("aria-selected")
);
```

你的第一個 `button` 元素應該有一個 `aria-selected` 屬性設為 `"true"`。

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-selected"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
