---
id: 68599633407374c0d574587b
title: 第 3 步
challengeType: 0
dashedName: step-3
---

# --description--

下一步是為你的按鈕添加負責在不同行星及其事實之間切換的功能。

在 `tablist` 元素內，為你的三個 `button` 元素添加 `role` 屬性並設為 `"tab"`。第一個 `button` 元素應該有文字 `🌍 Earth`。第二個 `button` 應該有文字 `🪐 Saturn`。第三個 `button` 元素應該有文字 `🔴 Mars`。

# --hints--

你應該在 `tablist` 元素內有三個 `button` 元素。

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

你的每個 `button` 元素都應該有一個設定為 `"tab"` 的 `role` 屬性。

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

你的第一個 `button` 元素應該有文字 `🌍 Earth`。

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

你的第二個 `button` 元素應該有文字 `🪐 Saturn`。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

你的第三個 `button` 元素應該有文字 `🔴 Mars`。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
