---
id: 6859a1270ac255c571352189
title: 第 8 步
challengeType: 0
dashedName: step-8
---

# --description--

現在是時候為每個面板創建標記，這些面板將顯示關於每個行星的事實。

首先為你的頁面添加三個 `div` 元素，每個元素的 `role` 屬性設為 `"tabpanel"`。在每個 `div` 元素內，應該有一個 `p` 元素。

# --hints--

你應該有三個 `div` 元素，其 `role` 屬性設為 `"tabpanel"`。

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"]'), 3);
```

每個 `tabpanel` 元素內都應該有一個 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('div[role="tabpanel"] p'), 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      --fcc-editable-region--
      
      --fcc-editable-region--
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
