---
id: 68599f9df77791c4a311cc5f
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

為每個 `button` 元素添加的最後一個項目是 `id` 屬性。這些將用於將 `button` 元素與其對應的內容面板連結。

第一個 `button` 元素應該有一個 `id` 為 `"tab-earth"`，第二個應該有一個 `id` 為 `"tab-saturn"`，第三個應該有一個 `id` 為 `"tab-mars"`。

# --hints--

你的第一個 `button` 元素應該有一個 `id` 屬性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("id")
);
```

你的第一個 `button` 元素應該有一個 `id` 屬性設定為 `"tab-earth"`。

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("id"), "tab-earth");
```

你的第二個 `button` 元素應該有一個 `id` 屬性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[1].hasAttribute("id")
);
```

你的第二個 `button` 元素應該有一個 `id` 屬性設定為 `"tab-saturn"`。

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("id"), "tab-saturn");
```

你的第三個 `button` 元素應該有一個 `id` 屬性。

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[2].hasAttribute("id")
);
```

你的第三個 `button` 元素應該有一個 `id` 屬性設為 `"tab-mars"`。

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("id"), "tab-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth" aria-selected="true">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
