---
id: 6859ab9ba7a915c84fa69851
title: 第 11 步
challengeType: 0
dashedName: step-11
---

# --description--

對於第二個 `tabpanel` 元素，為其添加 `id` 屬性，值為 `"panel-saturn"`，以及 `aria-labelledby` 屬性，值為 `"tab-saturn"`。

對於第三個 `tabpanel` 元素，為它添加 `id` 屬性，值為 `"panel-mars"`，以及 `aria-labelledby` 屬性，值為 `"tab-mars"`。

# --hints--

你的第二個 `tabpanel` 元素應該有一個 `id` 屬性。

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[1].hasAttribute("id"));
```

你的第二個 `id` 屬性應該有值 `"panel-saturn"`。

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[1].getAttribute("id"), "panel-saturn");
```

你的第二個 `tabpanel` 元素應該有一個 `aria-labelledby` 屬性。

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[1].hasAttribute("aria-labelledby"));
```

你的第二個 `aria-labelledby` 屬性應該有值 `"tab-saturn"`。

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[1].getAttribute("aria-labelledby"), "tab-saturn");
``` 

你的第三個 `tabpanel` 元素應該有一個 `id` 屬性。

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[2].hasAttribute("id"));
```

你的第三個 `id` 屬性應該有值 `"panel-mars"`。

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[2].getAttribute("id"), "panel-mars");
```

你的第三個 `tabpanel` 元素應該有一個 `aria-labelledby` 屬性。

```js
assert.isTrue(document.querySelectorAll('[role="tabpanel"]')[2].hasAttribute("aria-labelledby"));
```

你的第三個 `aria-labelledby` 屬性應該有值 `"tab-mars"`。

```js
assert.equal(document.querySelectorAll('[role="tabpanel"]')[2].getAttribute("aria-labelledby"), "tab-mars");
``` 

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        <button role="tab" aria-controls="panel-earth" aria-selected="true" id="tab-earth">🌍 Earth</button>
        <button role="tab" aria-controls="panel-saturn" aria-selected="false" id="tab-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars" aria-selected="false" id="tab-mars">🔴 Mars</button>
      </div>

      <div id="panel-earth" role="tabpanel" aria-labelledby="tab-earth">
        <p>
          Earth is our home planet, known for its abundant water, diverse ecosystems, and life-supporting atmosphere. It's the only planet in the solar system known to harbor life.
        </p>
      </div>
      --fcc-editable-region--
      <div role="tabpanel">
        <p>
          Saturn is famous for its beautiful and extensive ring system made of ice and rock particles. It's a gas giant with dozens of moons orbiting it.
        </p>
      </div>
      <div role="tabpanel">
      --fcc-editable-region--
        <p>
          Mars, the red planet, is a rocky world with the tallest volcano and deepest canyon in the solar system. It's a key focus for exploration in the search for past or present life.
        </p>
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
