---
id: 685db0bd1280154ae85bcf9f
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將使用 HTML、CSS 和 JavaScript 建置一個可存取的行星標籤列表。已為你提供 HTML 範本和 CSS。

在這前幾個步驟中，你將建置剩餘的 `HTML` 以完成 `tablist`。

首先建立一個 `div` 元素，並將類別設為 `"tabs"`。在該 `div` 元素內，添加一個 `h2` 元素，並將 `id` 屬性設為 `"tabs-title"`，文字內容為 `Planets`。

# --hints--

你應該有一個 `div` 元素。

```js
assert.exists(document.querySelector("div"));
```

你的 `div` 元素應該設定類別為 `"tabs"`。

```js
assert.exists(document.querySelector("div.tabs"));
```

你的 `div` 元素內應該有一個 `h2` 元素。

```js
assert.exists(document.querySelector(".tabs h2"));
```

你的 `h2` 元素應該有一個 `id` 屬性設定為 `"tabs-title"`。

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

你的 `h2` 元素應該有文字內容 `Planets`。

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
