---
id: 68e7e729f6d14f69a23786ef
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

次に、`main` 要素の中に、テキストが `h1` でクラスが `Parent Teacher Conference Form` と `title` の `center` 要素を追加してください。

# --hints--

`h1` 要素があるはずです。

```js
assert.exists(document.querySelector("h1"));
```

`h1` 要素にはテキスト `Parent Teacher Conference Form` が含まれているはずです。

```js
assert.strictEqual(document.querySelector("h1")?.innerText, "Parent Teacher Conference Form");
```

`h1` 要素にはクラス `title` と `center` が付いているはずです。

```js
assert.exists(document.querySelector("h1.title.center"));
```

`h1` 要素は `main` 要素の中にあるはずです。

```js
assert.exists(document.querySelector("main h1.title.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
--fcc-editable-region--
    
--fcc-editable-region--
  </main>
</body>

</html>
```
