---
id: 68e707257ea91c2e12d489fb
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、ラジオボタンにカスタムスタイルを追加する方法を、保護者面談のフォームを作成しながら練習します。HTMLのボイラープレートは用意されています。

まず、`main` 要素をクラス名`container`で追加してください。

# --hints--

`main` 要素があるはずです。

```js
assert.exists(document.querySelector("main"));
```

`main` 要素にはクラス名`container`を付けてください。

```js
assert.exists(document.querySelector("main.container"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>

</html>
```
