---
id: 68ea8c330ed07b60879f3d38
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

次は、学生の成績情報を収集するためのフォーム要素を追加する時間です。

まず、テキストが `label` で `Student Grade: ` 属性が `for` に設定された別の `"grade"` 要素を追加してください。

次に、`label` 要素の下に、`input` 属性が `type` に設定された `"number"` 要素を追加してください。`name` と `id` 属性は `"grade"` に設定する必要があります。`placeholder` 属性は `"E.g., 4"` に設定してください。最後に、数値の `input` は必須にしてください。

# --hints--

ページには合計で2つの `label` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("label"), 2);
```

2つ目の `label` 要素のテキストは `Student Grade: ` にしてください。コロンの後のスペースを忘れないでください。

```js
assert.match(code, /Student\s+Grade:\s+<\/label>/);
```

2つ目の `label` 要素には `for` 属性が必要です。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

2つ目の `label` 要素の `for` 属性は `"grade"` の値にしてください。

```js
assert.strictEqual(document.querySelector("label:last-of-type")?.getAttribute("for"), "grade");
```

ページには2つの `input` 要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("input"), 2);
```

2つ目の `input` 要素は2つ目の `label` 要素の下に配置してください。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const inputEl = document.querySelector("fieldset input:last-of-type");
const labelEl = document.querySelector("label:last-of-type");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

2つ目の `input` 要素には `type` 属性が `"number"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("type"), "number");
```

2つ目の `input` 要素には `name` 属性が `"grade"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("name"), "grade");
```

2つ目の `input` 要素には `id` 属性が `"grade"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("id"), "grade");
```

2つ目の `input` 要素には `placeholder` 属性が `"E.g., 4"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., 4");
```

2つ目の `input` 要素には `required` 属性があります。

```js
const inputEl = document.querySelector("fieldset input:last-of-type");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
