---
id: 68ea7c41f2555e16174e7cfa
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

次に、テキストが `Please fill out the form below to help schedule your parent-teacher conference.` の段落要素を追加してください。段落要素には `description` と `center` のクラスも付けてください。

# --hints--

`p` 要素があるはずです。

```js
assert.exists(document.querySelector("p"));
```

`p` 要素にはテキスト `Please fill out the form below to help schedule your parent-teacher conference.` が含まれているはずです。

```js
assert.strictEqual(document.querySelector("p")?.innerText, "Please fill out the form below to help schedule your parent-teacher conference.");
```

`p` 要素には `description` と `center` のクラスが付いているはずです。

```js
assert.exists(document.querySelector("p.description.center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
