---
id: 68ea81da9f687732368a5322
title: ステップ 4
challengeType: 0
dashedName: step-4
---

# --description--

次は、親と生徒の情報を表す `form` と `input` 要素を追加するステップです。

`form` 要素の下に `p` 要素を追加することから始めてください。

# --hints--

`form` 要素があるはずです。

```js
assert.exists(document.querySelector(".container form"));
```

form 要素は `p` 要素の下に配置してください。

```js
assert.exists(document.querySelector(".container > p+form"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>
    --fcc-editable-region--
    
    --fcc-editable-region--
  </main>
</body>

</html>
```
