---
id: 68eaa668d4d9fbe529a0e16b
title: ステップ 14
challengeType: 0
dashedName: step-14
---

# --description--

今、テキストが`label`で、`Any specific concerns or topics you'd like to discuss?`属性が`for`に設定された`notes`要素を追加してください。

この`label`の下に`textarea`要素を追加します。`id`を`notes`に、`name`を`notes`に、`rows`を`4`に、`cols`を`50`に設定してください。

# --hints--

最後の`fieldset`には`label`要素が含まれている必要があります。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

`label`要素はテキストが`Any specific concerns or topics you'd like to discuss?`である必要があります。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.innerText, "Any specific concerns or topics you'd like to discuss?");
```

`label`要素の`for`属性は`"notes"`に設定してください。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "notes");
```

最後の`fieldset`には`textarea`要素が含まれている必要があります。

```js
assert.exists(document.querySelector("fieldset:last-of-type textarea"));
```

`textarea`要素は`label`要素の下にある必要があります。

```js
const label = document.querySelector("fieldset:last-of-type label");
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(label?.nextElementSibling, textarea);
```

`textarea`要素は`id`属性が`"notes"`に設定されている必要があります。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("id"), "notes");
```

`textarea`要素は`name`属性が`"notes"`に設定されている必要があります。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("name"), "notes");
```

`textarea`要素は`rows`属性が`4`に設定されている必要があります。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("rows"), "4");
```

`textarea`要素は`cols`属性が`50`に設定されている必要があります。

```js
const textarea = document.querySelector("fieldset:last-of-type textarea");
assert.strictEqual(textarea?.getAttribute("cols"), "50");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
