---
id: 68ea975c7cbbb894a386e402
title: ステップ 11
challengeType: 0
dashedName: step-11
---

# --description--

次のステップは、メール連絡方法用の `label` と `input` 要素を追加することです。

まず、`"contact-method"` クラスを持ち、`for` 属性が `"email"` に設定されたラベルを追加してください。ラベルのテキストは `Email: ` にしてください。

`label` 要素の下に、`id` と `value` 属性が `"email"` に設定されたラジオボタンを追加してください。`name` 属性は `"contact-method"` に設定し、クラスは `"contact-method-radio-btn"` にしてください。

最後に、このラジオボタンがデフォルトでチェックされていることを確認してください。

# --hints--


最後のフィールドセットには `label` 要素が含まれているはずです。

```js
assert.exists(document.querySelector("fieldset:last-of-type label"));
```

この新しい `label` 要素（メール用）は `"contact-method"` クラスを持つべきです。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

新しい `label` 要素の `for` 属性は `"email"` に設定されているべきです。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.strictEqual(label?.getAttribute("for"), "email");
```

`label` 要素には `Email: ` というテキストがあるべきです。コロンの後のスペースも忘れずに。

```js
const label = document.querySelector("fieldset:last-of-type label");
assert.match(label?.innerText, /Email:\s?/);
```

最後の `input` の中に `fieldset` 要素があるはずです。

```js
assert.exists(document.querySelector("fieldset:last-of-type input"));
```

`input` 要素には `type` 属性が `radio` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 要素には `id` 属性が `"email"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("id"), "email");
```

`input` 要素には `value` 属性が `"email"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("value"), "email");
```

`input` 要素には `name` 属性が `"contact-method"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 要素は `"contact-method-radio-btn"` クラスを持つべきです。

```js
assert.exists(document.querySelector("fieldset:last-of-type input.contact-method-radio-btn"));
```

`radio` ボタンはデフォルトで `checked` であるべきです。

```js
const input = document.querySelector("fieldset:last-of-type input");
assert.isTrue(input?.hasAttribute("checked"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
