---
id: 68ea8a5db73b155708992687
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

次のステップは、学生の情報に関連する `input` 要素を追加することです。

まず、`input` 属性を `type` に設定した `"text"` 要素を追加してください。次に、`name` と `id` 属性を両方とも `student-name` に設定します。

次に、`placeholder` 属性を `E.g., Jane Doe` に設定してください。最後に、`required` 属性を追加します。

# --hints--

`input` 要素があるはずです。

```js
assert.exists(document.querySelector("fieldset input"));
```

`input` 要素には `type` 属性が `"text"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

`input` 要素には `name` 属性が `"student-name"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("name"), "student-name");
```

`input` 要素には `id` 属性が `"student-name"` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("id"), "student-name");
```

`input` 要素には `placeholder` 属性が `E.g., Jane Doe` に設定されている必要があります。

```js
const inputEl = document.querySelector("fieldset input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Jane Doe");
```

`input` 要素には `required` 属性が必要です。

```js
const inputEl = document.querySelector("fieldset input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>

      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>


    </form>
  </main>
</body>

</html>
```
