---
id: 68ea8fdf2c9b167307e7f671
title: ステップ 9
challengeType: 0
dashedName: step-9
---

# --description--

次は、親の情報を収集するためのフォーム要素を追加する時間です。

まず、テキストが`label`で、`Parent/Guardian Name: `属性が`for`に設定された`"parent-name"`要素を追加してください。

次に、`label`要素の下に、`input`属性が`type`に設定された`"text"`要素を追加してください。`name`と`id`属性は`"parent-name"`に設定する必要があります。`placeholder`属性は`"E.g., Nancy Doe"`に設定してください。最後に、`input`は必須にしてください。

# --hints--

ページには合計で3つの`label`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("label"), 3);
```

3つ目の`label`要素のテキストは`Parent/Guardian Name: `であるべきです。コロンの後のスペースを忘れないでください。

```js
assert.match(code, /Parent\/Guardian\s+Name:\s+<\/label>/);
```

3つ目の`label`要素には`for`属性が必要です。

```js
assert.isTrue(document.querySelector("label:last-of-type")?.hasAttribute("for"));
```

3つ目の`label`要素の`for`属性は`"parent-name"`の値であるべきです。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type label")?.getAttribute("for"), "parent-name");
```

ページには3つの`input`要素があるはずです。

```js
assert.lengthOf(document.querySelectorAll("input"), 3);
```

3つ目の`input`要素は3つ目の`label`要素の下にあるべきです。

```js
const fieldsetChildren = [...document.querySelector("fieldset:last-of-type")?.children];
const inputEl = document.querySelector("fieldset:last-of-type input");
const labelEl = document.querySelector("fieldset:last-of-type label");

assert.isBelow(fieldsetChildren.indexOf(labelEl), fieldsetChildren.indexOf(inputEl));
```

3つ目の`input`要素には`type`属性が`"text"`に設定されているべきです。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("type"), "text");
```

3つ目の`input`要素には`name`属性が`"parent-name"`に設定されているべきです。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("name"), "parent-name");
```

3つ目の`input`要素には`id`属性が`"parent-name"`に設定されているべきです。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("id"), "parent-name");
```

3つ目の`input`要素には`placeholder`属性が`"E.g., Nancy Doe"`に設定されているべきです。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.strictEqual(inputEl?.getAttribute("placeholder"), "E.g., Nancy Doe");
```

3つ目の`input`要素には`required`属性が必要です。

```js
const inputEl = document.querySelector("fieldset:last-of-type input");
assert.isTrue(inputEl?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
