---
id: 68ea8666b8b473455f5b33ef
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

フォームの最初のセクションは、学生の名前や学年などの情報にフォーカスします。

`form` 要素の中に `fieldset` 要素を追加してください。その `fieldset` 要素の中に、テキスト `legend` を持つ `Student Information` 要素を追加してください。

`legend` 要素の下に、テキスト `label` を持ち、`Full Name: ` の値を持つ `for` 属性を持つ `student-name` 要素を追加してください。


# --hints--

`fieldset` 要素があるはずです。

```js
assert.exists(document.querySelector("fieldset"));
```

`fieldset` 要素は `form` 要素の中にあるはずです。

```js
assert.exists(document.querySelector("form fieldset"));
```

`legend` 要素があるはずです。

```js
assert.exists(document.querySelector("legend"));
```

`legend` 要素は `fieldset` 要素の中にあるはずです。

```js
assert.exists(document.querySelector("fieldset legend"));
```

`legend` 要素はテキスト `Student Information` を持つはずです。

```js
assert.strictEqual(document.querySelector("fieldset legend")?.innerText, "Student Information");
```

`label` 要素があるはずです。

```js
assert.exists(document.querySelector("label"));
```

`label` 要素はテキスト `Full Name: ` を持つはずです。コロンの後のスペースを忘れないでください。

```js
assert.match(code, /Full\s+Name:\s+<\/label>/);
```

`label` 要素は `for` 属性を持つはずです。

```js
assert.isTrue(document.querySelector("label")?.hasAttribute("for"));
```

`label` 要素の `for` 属性は `student-name` の値を持つはずです。

```js
assert.strictEqual(document.querySelector("label")?.getAttribute("for"), "student-name");
```

`label` 要素は `legend` 要素の下にあるはずです。

```js
const fieldsetChildren = [...document.querySelector("fieldset")?.children];
const legendEl = document.querySelector("fieldset legend");
const labelEl = document.querySelector("label");

assert.isBelow(fieldsetChildren.indexOf(legendEl), fieldsetChildren.indexOf(labelEl));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
