---
id: 68ea9f01c51b66b937793611
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

次に、クラスが `"contact-method"` で、`for` 属性が `"phone"` に設定された別のラベルを追加してください。ラベルのテキストは `Phone: ` にしてください。

`label` 要素の下に、`id` と `value` 属性が `"phone"` に設定されたラジオボタンを追加してください。`name` 属性は `"contact-method"` に設定し、クラスは `"contact-method-radio-btn"` に設定してください。

# --hints--

最後の `label` の中に `fieldset` 要素が2つあるはずです。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type label"), 2);
```

`label` 要素はクラスが `"contact-method"` であるべきです。

```js
assert.exists(document.querySelector("fieldset:last-of-type label.contact-method"));
```

`label` 要素の `for` 属性は `"phone"` に設定されているべきです。

```js
assert.exists(document.querySelector("fieldset:last-of-type label[for='phone']"));
```

`label` 要素のテキストは `Phone: ` であるべきです。

```js
const label = document.querySelector("fieldset:last-of-type label[for='phone']");
assert.match(label?.innerText, /Phone:\s?/);
```

最後の `input` の中に `fieldset` 要素が2つあるはずです。

```js
assert.lengthOf(document.querySelectorAll("fieldset:last-of-type input"), 2);
```

`input` 要素は `type` 属性が `radio` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("type"), "radio");
```

`input` 要素は `id` 属性が `"phone"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("id"), "phone");
```

`input` 要素は `value` 属性が `"phone"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("value"), "phone");
```

`input` 要素は `name` 属性が `"contact-method"` に設定されているべきです。

```js
const input = document.querySelector("fieldset:last-of-type input:last-of-type");
assert.strictEqual(input?.getAttribute("name"), "contact-method");
```

`input` 要素はクラスが `"contact-method-radio-btn"` であるべきです。

```js
assert.exists(document.querySelector("fieldset:last-of-type input[id='phone'].contact-method-radio-btn"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />
        --fcc-editable-region--
        
        --fcc-editable-region--
      </fieldset>
    </form>
  </main>
</body>

</html>
```
