---
id: 68ea8e6a87a2816b0bb2eb7d
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

次のステップは、親情報のためのフォームのセクションです。

まず、もう一つの`fieldset`要素を追加してください。その`fieldset`要素の中に、テキスト`legend`を持つ`Parent/Guardian Information`要素を追加してください。

# --hints--

ページには合計で2つの`fieldset`要素があります。

```js
assert.lengthOf(document.querySelectorAll("fieldset"), 2);
```

2番目の`legend`要素の中に`fieldset`要素があるはずです。

```js
assert.exists(document.querySelector("fieldset:last-of-type legend"));
```

`legend`要素にはテキスト`Parent/Guardian Information`があるはずです。

```js
assert.strictEqual(document.querySelector("fieldset:last-of-type legend")?.innerText, "Parent/Guardian Information");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />


        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </form>
  </main>
</body>

</html>
```
