---
id: 68eb058081504963447bfa60
title: ステップ 17
challengeType: 0
dashedName: step-17
---

# --description--

次に、メインのコンテナをスタイル設定しましょう。クラスが `container` の要素を選択し、その `background-color` を16進コード `#ffffff1a` に設定してください。

16進コードは、HTMLやCSSで色を表すために使われる数字と文字の6桁の組み合わせです。先頭は `#` で始まり、続いて3組のペアが並びます。

- 最初の2桁は赤を表します
- 次の2桁は緑を表します
- 最後の2桁は青を表します

例えば、`#ffffff` は赤、緑、青の値が最大なので純粋な白色です。

また、末尾に2桁追加して不透明度（アルファチャネルと呼ばれます）を調整することもできます。`#ffffff1a` では `1a` によって白色が半透明になっています。アルファ値が低いほど色は透明に見えます。16進コードについては後でさらに学びます！

# --hints--

クラスが `container` の要素を対象にしてください。

```js
assert(new __helpers.CSSHelp(document).getStyle('.container'));
```

`background-color` プロパティを `#ffffff1a` に設定してください。

```js
const style = new __helpers.CSSHelp(document).getStyle('.container')
assert.strictEqual(style?.backgroundColor, 'rgba(255, 255, 255, 0.1)')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Parent Teacher Conference Form</title>
  <link rel="stylesheet" href="./styles.css" />
</head>

<body>
  <main class="container">
    <h1 class="title center">Parent Teacher Conference Form</h1>
    <p class="description center">Please fill out the form below to help schedule your parent-teacher conference.</p>

    <form>
      <fieldset>
        <legend>Student Information</legend>
        <label for="student-name">Full Name: </label>
        <input
            type="text"
            name="student-name"
            id="student-name"
            required
            placeholder="E.g., Jane Doe"
          />

        <label for="grade">Student Grade: </label>
        <input
            type="number"
            name="grade"
            id="grade"
            required
            placeholder="E.g., 4"
          />
      </fieldset>

      <fieldset>
        <legend>Parent/Guardian Information</legend>
        <label for="parent-name">Parent/Guardian Name: </label>
        <input
            type="text"
            name="parent-name"
            id="parent-name"
            required
            placeholder="E.g., Nancy Doe"
          />
      </fieldset>

      <fieldset>
        <legend>Preferred Contact Method</legend>
        <label class="contact-method" for="email">Email: </label>
        <input
            id="email"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="email"
            checked
          />

        <label class="contact-method" for="phone">Phone: </label>
        <input
            id="phone"
            class="contact-method-radio-btn"
            type="radio"
            name="contact-method"
            value="phone"
          />
      </fieldset>

      <fieldset>
        <legend>Additional Notes</legend>
        <label for="notes"
            >Any specific concerns or topics you'd like to discuss?</label>
        <textarea id="notes" name="notes" rows="4" cols="50"></textarea>
      </fieldset>

      <button class="submit-btn" type="submit">Submit Form</button>
    </form>
  </main>
</body>

</html>
```


```css
body {
  background-color: MidnightBlue;
  color: whitesmoke;
}

--fcc-editable-region--

--fcc-editable-region--
```
