---
id: 68de7360e1832ea180577a8a
title: 构建一个第 N 个斐波那契数计算器
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `fibonacci` 的函数。
2. 你应该在 `fibonacci` 函数中定义一个名为 `sequence` 的列表，并且它应初始化为值 `[0, 1]`。
3. `fibonacci` 函数应接受一个参数，一个正整数 `n`。
4. 调用 `fibonacci(n)` 应该使用动态编程方法来计算并返回斐波那契序列中的第 `n` 个数字，其中每个数字是前述的两个数字之和。
5. 斐波那契序列中位置为 `n` 的每个计算数字应保存在索引为 `n - 1` 的 `sequence` 列表中。

# --hints--

你应该有一个名为 `fibonacci` 的函数。

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

你的 `fibonacci` 函数应接受一个参数。

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

你应该在 `fibonacci` 函数内有一个名为 `sequence` 的列表，初始化为 `[0, 1]`。

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` 应该返回 `0`。

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` 应该返回 `1`。

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` 应该返回 `1`。

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` 应该返回 `2`。

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` 应该返回 `5`。

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` 应该返回 `55`。

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` 应该返回 `610`。

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

你不应该在你的代码中使用递归。

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

