---
id: 6551f6f9ec6db270eb83f3db
title: 步驟 8
challengeType: 20
dashedName: step-8
---

# --description--

您可以在終端上看到數字 `11` 被打印出來，因爲 `'Hello World'` 包含11個字符。

另一個有用的內置函數是 `type()`, 它返回一個變量的數據類型。 修改您的代碼`print()`來調用並打印`text`的數據類型.

# --hints--

You should call the `type()` function.

```js
assert.match(code, /type\s*\(.*\)/)
```

You should pass `text` to the `type()` function by including it between the parentheses.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

You should pass `type(text)` to the `print()` function by including it between the parentheses. 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

You should not have `print(len(text))` in your code.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
