---
id: 6551f3bfed58796ebb1268dc
title: 步驟5
challengeType: 20
dashedName: step-5
---

# --description--

每個字符串都可以被一個數字索引引用。 索引數開始於零。 所以字符串的第一個字符有一個 `0` 的索引。 例如，在字符串 `'Hello World'`, `'H'` 位於索引 `0`, `'e'` 位於索引 `1`等。

字符串的每個字符都可以使用括號符號訪問。 您需要在方括號前寫上變量名稱並在方括號中添加字符索引：

```py
text = 'Hello World'
r = text[8]
```

現在，不是打印 `文本`，而是隻打印索引 `6` 號位的字符。

# --hints--

You should still call the `print()` function.

```js
assert.match(code, /print\s*\(.*\)/)
```

You should pass `text[6]` to the `print()` function by including it between the parentheses. 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
