---
id: 6551f4259d66026ee28cb205
title: 步驟6
challengeType: 20
dashedName: step-6
---

# --description--

您也可以從字符串結尾處開始訪問字符。 最後一個字符的索引爲 `-1`, 倒數第二個是 `-2` ，以此類推。

現在修改您現有的 `print()` 來調用並打印你字符串裏的最後一個字符。

# --hints--

You should still call the `print()` function.

```js
assert.match(code, /print\s*\(.*\)/)
```

You should pass `text[-1]` to the `print()` function by including it between the parentheses. Pay attention to place the function call at the beginning of the line.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

您的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
