---
id: 6551ecf597038c36637f3759
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

變量是Python和任何編程語言的重要部分。 變量是一個引用某對象或指向某對象的名稱。 您可以在賦值運算符`=`的左側寫入變量的名稱以設一個變量，並在賦值運算符右側指定分配給該變量的值。

```py
variable_name = value
```

請創建一個名爲`number`的變量，並分配`5`爲該變量的值。

# --hints--

你應設一個名爲`number`的變量。 請注意，需把變量的名稱放在該行的開頭位置。

```js
assert.match(code, /^number\s*=/m)
```

您應將`5`分配爲變量`number`。

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

您的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
