---
id: 6551fe3b1df7c9740f13f270
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

Python變量命名的關鍵是：

- 有些詞是系統保留的關鍵詞(例如 `for`, `while`, `True`)。 他們在 Python 語言裏有特殊的含義，所以你不能將它們用作變量名稱。
- 變量名稱不能以數字開頭，它們只能包含字母、數字和下劃線。
- 變量名稱是區分大小寫的，即 `my_var` 不同於 `my_Var` 和 `MY_VAR`
- 最後，用 `snake_case`寫變量名稱是一個慣例， 其中每個空格被下劃線字符所取代，單詞用小寫字母寫出。

刪除兩個 `print()` ,並定義另一個變量名爲 `alphabet`。 Assign the string `'abcdefghijklmnopqrstuvwxyz'` to this variable.

# --hints--

You should not have `print(type(text))` in your code.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

You should not have `print(type(shift))` in your code.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

You should declare a variable called `alphabet`. Pay attention to place the variable name at the beginning of the line.

```js
assert.match(code, /^alphabet\s*=/m)
```

You should assign the string `'abcdefghijklmnopqrstuvwxyz'` to your `alphabet` variable. Remember to use either single or double quotes to enclose the string.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
