---
id: 6551ef504d91936d2d4e54f8
title: 步驟4
challengeType: 20
dashedName: step-4
---

# --description--

*argument*是傳遞給函數的對象或表達式 - 在調用時添加在左括號和右括號之間：

```py
greet = 'Hello!'
print(greet)
```

上面例子中的代碼將打印字符串 `'Hello!'`，它是作爲參數傳遞給 `print()` 的變量 `greet` 的值。

通過將 `text` 變量作爲參數傳遞給 `print()` 函數，將 `text` 變量打印到屏幕上。

# --hints--

You should pass `text` to the `print()` function by adding the name of this variable within the opening and closing parentheses. 注意將函數調用放在行首。

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

您的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
