---
id: 6552008c0d9d9075cbec9772
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

你將使用 `.find()` 方法來查找消息中每個字母在字母表中的位置。 方法類似於函數，但它屬於對象。

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Above, the `.find()` method is *called on* `sentence` (the string to search in), and `'r'` (the character to locate) is passed as the argument. The `sentence.find('r')` call will return `4`, which is the index of the first occurrence of `'r'` in `sentence`.

At the end of your code, call `.find()` on `alphabet` and pass `'z'` as the argument to the method.

# --hints--

你應該調用 `.find()` 方法。

```js
assert.match(code, /find\s*\(.*\)/)
```

你應該對 `alphabet` 變量調用 `.find()` 方法。

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


你應該對 `alphabet` 變量調用 `.find()` 方法，並將 `'z'` 傳遞給該方法。 注意將函數調用放在行首。

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


你的代碼中包含無效的句法和/或無效的縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
