---
id: 65521203d7165c7b84b22ad4
title: 步驟-23
challengeType: 20
dashedName: step-23
---

# --description--

A loop allows you to systematically go through a sequence of elements and execute actions on each one.

In this case, you'll employ a `for` loop. 以下是迭代 `text` 的方法：

```py
for i in text:
```

`for` 是表示循環類型的關鍵字。 `i` 是一個變量，它按順序獲取 `text` 中元素的值。 該語句以冒號 `:` 結尾。

Below the line where you declared `alphabet`, write a `for` loop to iterate over `text`. 使用 `i` 作爲循環變量。

這樣做時，終端出現錯誤。 你們將在下一步驟中學習這個問題。

# --hints--

你應該使用 `for` 關鍵字來創建循環。 確保將 `for` 關鍵字放在行的開頭，並在關鍵字後面留一個空格。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

你應該在 `for` 關鍵字後面寫入 `i` 變量

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

你應該在 `for i` 之後寫入 `in` 關鍵字。 確保在 `in` 關鍵字周圍留一個空格

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

你應該在 `for i in` 之後寫入 `text`。 不要忘記添加最後的 `:`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

你的 `for` 循環應放置在代碼 `alphabet = 'abcdefghijklmnopqrstuvwxyz'` 行下方。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
