---
id: 6553995f412dd8122ed38e4a
title: 第 43 步
challengeType: 20
dashedName: step-43
---

# --description--

條件語句還可以有 `else` 子句。 可以將此子句添加到 `if` 語句的末尾，以便在 `if` 語句的條件爲假時執行替代代碼：

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

As you can see in your output, when the loop iterations reach the space, a space is added to the encrypted string. 然後執行 `if` 塊外部的代碼，並將 `c` 添加到加密字符串中。

要修復此問題，請在 `encrypted_text += char` 之後添加 `else` 子句，並縮進除 `print()` 調用之外的所有後續代碼行。

# --hints--

You should create an `else` clause. 請記住包含最後的冒號。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

You should indent the lines of code after your `else` clause except the `print()` call.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
