---
id: 6553a7d8d05cbb1ae335a665
title: 步驟 45
challengeType: 20
dashedName: step-45
---

# --description--

當循環到達字母 `Z` 時，`index + shift` 之和超過了字符串 `alphabet` 的最後一個索引。 因此，`alphabet[new_index]` 嘗試使用無效索引，這會導致拋出 `IndexError` 異常。

You can notice that the output in the terminal stops at the space immediately before the `Z`, the last `print` before the error is thrown.

在這種情況下，模運算符 (`%`) 可用於返回兩個數字相除的餘數。 例如：`5 % 2` 等於 `1`，因爲 5 除以 2 的商爲 2，餘數爲 1。

將 `index + shift` 用括號括起來，並用 `26`（即 `alphabet` 長度）對錶達式取模。

# --hints--

You should have `new_index = (index + shift) % 26` in your `else` statement.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Your code contains invalid syntax and/or invalid indentation.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
