---
id: 6553f8c570f9982e013a8886
title: 步驟 57
challengeType: 20
dashedName: step-57
---

# --description--

目前，根據指定的偏移量，每個字母總是用相同的字母加密。 如果每個字母的偏移量都不同呢？ 這樣就更難解密了。 該算法稱爲維吉尼亞密碼，其中每個字母的偏移量由另一個文本（稱爲*key*）確定。

通過刪除兩個函數調用，開始將凱撒密碼轉換爲維吉尼亞密碼

# --hints--

You should remove your two `caesar()` function calls.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
