---
id: 6554a334a40edb1fb4eff827
title: 步驟 62
challengeType: 20
dashedName: step-62
---

# --description--

接下來，在 `else` 塊內，聲明一個名爲 `key_char` 的變量，並將索引 `key_index` 處的 `key` 值 mod(`%`) `key` 的長度分配給它。

# --hints--

你應該在 `else` 子句主體的開頭聲明一個名爲 `key_char` 的變量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/^\s*key_char\s*=/));
```

你應該將 `key[key_index % len(key)]` 分配給 `key_char` 變量。

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");    
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]/));
```

你的代碼包含無效語法和/或無效縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--        
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
--fcc-editable-region--        
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
