---
id: 659db3e3670d3e712be82593
title: 步驟 22
challengeType: 20
dashedName: step-22
---

# --description--

重複在字母表中定位字母並確定 `text` 中每個字符的移位字母的過程可能很乏味。 值得慶幸的是，你可以使用循環來簡化它。

現在，移除所有在聲明 `alphabet` 變量之後的代碼行。

# --hints--

你的代碼中仍應包含 `text = 'Hello World'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

你的代碼中仍應包含 `shift = 3`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

你的代碼中仍應包含 `alphabet = 'abcdefghijklmnopqrstuvwxyz'`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

你應該刪除 `index` 變量及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

你的代碼中不應該有 `print(index)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

你應該刪除 `shifted` 變量及其值。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

你的代碼中不應該有 `print(shifted)`。

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

你的代碼包含無效的語法和/或無效縮進。

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
