---
id: 5f332b23c2045fb843337579
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

To let visitors know the cafe was founded in `2020`, add a `p` element below the `h1` element with the text `Est. 2020`.

# --hints--

Dovresti avere un tag di apertura `<p>`.

```js
assert.match(code, /<p>/i);
```

Dovresti avere un tag di chiusura `</p>`.

```js
assert.match(code, /<\/p>/i);
```

Non dovresti modificare l'elemento `h1` già esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

L'elemento `p` dovrebbe trovarsi sotto l'elemento `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Il tuo elemento `p` dovrebbe avere il testo `Est. 2020`.

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

