---
id: 5f332b23c2045fb843337579
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Per far sapere ai visitatori che il caffè è stato fondato in `2020`, aggiungi un elemento `p` sotto l'elemento `h1` con il testo `Est. 2020`.

# --hints--

Dovresti avere un tag di apertura `<p>`.

```js
assert.match(code, /<p>/i);
```

Dovresti avere un tag di chiusura `</p>`.

```js
assert.match(code, /<\/p>/i);
```

Non dovresti modificare il tuo elemento `h1` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Il tuo elemento `p` dovrebbe essere sotto il tuo elemento `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Il tuo elemento `p` dovrebbe contenere il testo `Est. 2020`.

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
