---
id: 5f33294a6af5e9188dbdb8f3
title: Passaggio 1
challengeType: 0
dashedName: step-1
---

# --description--

In questo workshop, eserciterai le basi del CSS (Cascading Style Sheets) costruendo un menu per un caffè.

Iniziamo aggiungendo un po' di contenuto al menu. Aggiungi un elemento `main` all'interno dell'elemento `body` esistente. Conterrà infine informazioni sui prezzi di caffè e dessert offerti dal caffè.

# --hints--

Il tuo codice dovrebbe avere un tag di apertura `<main>`.

```js
assert.match(code, /<main>/i);
```

Il tuo codice dovrebbe avere un tag di chiusura `</main>`.

```js
assert.match(code, /<\/main>/i);
```

Non dovresti modificare il tuo elemento `body`. Assicurati di non eliminare accidentalmente il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Il tuo tag `main` dovrebbe essere all'interno del tag `body`.

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
