---
id: 5f332a88dc25a0fd25c7687a
title: Passaggio 2
challengeType: 0
dashedName: step-2
---

# --description--

Il nome del caffè è `CAMPER CAFE`. Quindi, aggiungi un elemento `h1` all’interno del tuo elemento `main`. Dagli il nome del caffè in lettere maiuscole per farlo risaltare.

# --hints--

Dovresti avere un tag `<h1>` di apertura.

```js
assert.match(code, /<h1>/i);
```

Dovresti avere un tag `</h1>` di chiusura.

```js
assert.match(code, /<\/h1>/i);
```

Non dovresti modificare il tuo elemento `main` esistente.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Il tuo elemento `h1` dovrebbe essere annidato nel tuo elemento `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Il tuo elemento `h1` dovrebbe contenere il testo `CAMPER CAFE`.

```js
assert.equal(document.querySelector("h1")?.innerText.trim(), "CAMPER CAFE");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
