---
id: 5f344fc1520b6719f2e35605
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

Ci saranno due sezioni nel menu, una per i caffè e una per i dolci. Aggiungi un elemento `section` all’interno dell’elemento `main` in modo da avere un posto dove mettere tutti i caffè disponibili.

# --hints--

Dovresti avere un tag di apertura `<section>`.

```js
assert.match(code, /<section\s*>/i);
```

Dovresti avere un tag di chiusura `</section>`.

```js
assert.match(code, /<\/section\s*>/i);
```

Non dovresti modificare il tuo elemento `main` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

Il tuo elemento `section` dovrebbe essere all’interno del tuo elemento `main`.

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
