---
id: 5f344fad8bf01691e71a30eb
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Finora hai avuto un controllo limitato sulla presentazione e l'aspetto del tuo contenuto. Per modificarlo, aggiungi un elemento `style` all'interno dell'elemento `head`.

# --hints--

Il tuo codice dovrebbe avere un tag di apertura `<style>`.

```js
assert.match(code, /<style\s*>/i);
```

Il tuo codice dovrebbe avere un tag di chiusura `</style>`.

```js
assert.match(code, /<\/style\s*>/);
```

Il tuo elemento `style` dovrebbe essere annidato nel tuo elemento `head`.

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
