---
id: 5f3477cbcb6ba47918c1da92
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Per far sì che lo stile della pagina appaia simile su dispositivi mobili come su desktop o laptop, devi aggiungere un elemento `meta` con un attributo `content` speciale.

Hai già imparato l'elemento `meta` viewport nelle lezioni precedenti.

Ecco un esempio:

```html
<meta name="viewport" content="width=device-width, initial-scale=1.0" />
```

# --hints--

Il tuo codice dovrebbe contenere due elementi `meta`.

```js
assert.lengthOf(code.match(/<meta.*?\/?>/g),2);
```

Il tuo elemento `meta` dovrebbe avere un attributo `name` con valore `viewport`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/name=('|")viewport\1/) || meta[1]?.outerHTML?.match(/name=('|")viewport\1/));
```

Il tuo elemento `meta` dovrebbe avere un attributo `content` con valore `width=device-width, initial-scale=1.0`.

```js
const meta = document.querySelectorAll('meta');
assert.exists(meta[0]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/) || meta[1]?.outerHTML?.match(/content=('|")width=device-width, initial-scale=1.0\1/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
--fcc-editable-region--
    
--fcc-editable-region--
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
