---
id: 5f3477cb2e27333b1ab2b955
title: Passaggio 12
challengeType: 0
dashedName: step-12
---

# --description--

Ora devi collegare il file `styles.css` per applicare di nuovo gli stili. All’interno dell’elemento `head`, aggiungi un elemento `link`. Dagli un attributo `rel` con valore `"stylesheet"` e un attributo `href` con valore `"styles.css"`.

**Nota:** L’elemento `link` è un elemento void, cioè non ha un tag di chiusura. Gli elementi void devono essere scritti come `<link>` invece che `<link></link>`.

# --hints--

Il tuo codice dovrebbe avere un elemento `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

L’elemento `link` è un elemento void, non dovrebbe avere un tag di chiusura `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Non dovresti modificare l’elemento `head` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Dovresti avere un solo elemento `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Il tuo elemento `link` dovrebbe essere all’interno del tuo elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Il tuo elemento `link` dovrebbe avere un attributo `rel` con valore `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Il tuo elemento `link` dovrebbe avere un attributo `href` con valore `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
