---
id: 5f3477aefa51bfc29327200b
title: Passaggio 10
challengeType: 0
dashedName: step-10
---

# --description--

Hai definito lo stile di tre elementi scrivendo CSS all'interno dei tag `style`. Questo funziona, ma dato che ci saranno molti altri stili, è meglio mettere tutti gli stili in un file separato e collegarlo.

Per te è stato creato un file `styles.css` separato. Puoi passare da un file all'altro usando le schede in alto nell'editor.

Inizia riscrivendo gli stili che hai creato nel file `styles.css`. Assicurati di escludere i tag di apertura e chiusura `style`.

# --hints--

Il tuo file `styles.css` dovrebbe avere il selettore di tipo `h1, h2, p`.

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```
 
Il tuo selettore dovrebbe impostare la proprietà `text-align` su `center`.

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

Il tuo codice non dovrebbe contenere i tag `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
