---
id: 5f3477cb2e27333b1ab2b955
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

Ora è necessario collegare il file `styles.css` in modo che gli stili vengano applicati di nuovo. All'interno dell'elemento `head`, aggiungi un elemento `link`. Dagli un attributo `rel` con il valore `"stylesheet"` e un attributo `href` con il valore `"styles.css"`.

**Note:** The `link` element is a void element, which means it doesn't have a closing tag. Void elements should be written as `<link>` rather than `<link></link>`.

# --hints--

Il codice dovrebbe avere un elemento `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

L'elemento `link` è un elemento vuoto, non dovrebbe avere un tag finale `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Non dovresti cambiare l'elemento `head` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Dovresti avere un elemento `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

L'elemento `link` dovrebbe trovarsi all'interno dell'elemento `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

L'elemento `link` dovrebbe avere un attributo `rel` con il valore `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

L'elemento `link` dovrebbe avere un attributo `href` con il valore `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
