---
id: 5f356ed6199b0cdef1d2be8f
title: Passaggio 22
challengeType: 0
dashedName: step-22
---

# --description--

Finora hai usato selettori di tipo e id per definire lo stile degli elementi. Tuttavia, è più comune usare un selettore diverso per definire lo stile dei tuoi elementi.

Hai imparato a lavorare con i selettori `class` in lezioni precedenti come questa:

```css
.class-name {
  styles
}
```

Trasforma il selettore `#menu` esistente in un selettore di classe sostituendo `#menu` con una classe chiamata `.menu`.

# --hints--

Dovresti avere un selettore di classe `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert.exists(hasMenu);
```

Non dovresti avere un selettore `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert.notExists(hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
--fcc-editable-region--
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
