---
id: 5f35e5c4321f818cdc4bed30
title: Passaggio 25
challengeType: 0
dashedName: step-25
---

# --description--

Ora che tutto sembra a posto, è il momento di iniziare ad aggiungere alcuni elementi del menu.

Aggiungi un elemento `article` vuoto sotto l'intestazione `Coffee`. Conterrà il gusto e il prezzo di ogni caffè che offri attualmente.

# --hints--

Dovresti avere un tag di apertura `<article>`.

```js
assert.match(code,/<article>/i);
```

Dovresti avere un tag di chiusura `</article>`.

```js
assert.match(code, /<\/article>/i);
```

Non dovresti modificare l'elemento `h2` esistente.

```js
assert.lengthOf(document.querySelectorAll('h2'), 1);
```

Il tuo elemento `article` dovrebbe venire dopo il tuo elemento `h2`.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
