---
id: 5f356ed63c7807a4f1e6d054
title: Passaggio 17
challengeType: 0
dashedName: step-17
---

# --description--

L'obiettivo ora è fare in modo che il `div` non occupi tutta la larghezza della pagina. La proprietà CSS `width` è perfetta per questo.

Puoi usare il selettore `id` per mirare a un elemento specifico con un attributo `id`.

Hai imparato a lavorare con il selettore `id` in lezioni precedenti come questa:

```css 
#cat {
  width: 250px;
}
```

Usa il selettore `#menu` per dare al tuo elemento una larghezza di `300px`.

# --hints--

Dovresti avere un selettore `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert.exists(hasDiv);
```

Dovresti impostare la proprietà `width` su `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isTrue(hasWidth);
```

Il tuo `div` dovrebbe avere una larghezza di 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert.equal(divWidth, '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
