---
id: 5f356ed60a5decd94ab66986
title: Passaggio 18
challengeType: 0
dashedName: step-18
---

# --description--

I commenti in CSS si scrivono così:

```css
/* comment here */
```

Nel tuo foglio di stile, commenta la riga che contiene la proprietà e il valore `background-color`, così potrai vedere l’effetto di definire lo stile solo dell’elemento `#menu`. Questo farà tornare lo sfondo bianco.

# --hints--

Dovresti commentare la riga `background-color: burlywood;` nel tuo CSS.

```js
assert.match(code, /\/\*\s*background-color:\s*burlywood;?\s*\*\//i);
```


Il tuo `body` dovrebbe avere uno sfondo bianco.

```js
const body = document.querySelector("body");
const compStyles = window.getComputedStyle(body);
const backgroundColor = compStyles?.getPropertyValue('background-color');
assert.equal(backgroundColor, "rgba(0, 0, 0, 0)")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
--fcc-editable-region--
  background-color: burlywood;
--fcc-editable-region--
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 300px;
}
```
