---
id: 5f3c866dbf362f99b9a0c6d0
title: Passaggio 33
challengeType: 0
dashedName: step-33
---

# --description--

Gli elementi `p` sono annidati in un elemento `article` con l'attributo classe `item`. Puoi definire lo stile di tutti gli elementi `p` annidati ovunque in elementi con una classe chiamata `item` in questo modo:

```css
.item p { }
```

Usando il selettore sopra, aggiungi una proprietà `display` con valore `inline-block` in modo che gli elementi `p` si comportino più come elementi in linea.

# --hints--

Dovresti usare il selettore `.item p`.

```js
const itemPStyle = new __helpers.CSSHelp(document).getStyle('.item p');
assert.exists(itemPStyle);
```

Dovresti impostare la proprietà `display` su `inline-block`.

```js
const hasDisplay = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.display === 'inline-block');
assert.isTrue(hasDisplay);
```

Il tuo selettore `.item p` dovrebbe impostare la proprietà `display` su `inline-block`.

```js
const itemPDisplay = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('display');
assert.equal(itemPDisplay, 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

--fcc-editable-region--

--fcc-editable-region--

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
