---
id: 5f3c866daec9a49519871816
title: Passaggio 26
challengeType: 0
dashedName: step-26
---

# --description--

Gli elementi `article` comunemente contengono più elementi che hanno informazioni correlate. In questo caso, conterrà un gusto di caffè e un prezzo per quel gusto.

Annida due elementi `p` all'interno del tuo elemento `article`. Il testo del primo deve essere `French Vanilla`, e il testo del secondo `3.00`.

# --hints--

Non dovresti modificare l'elemento `article` esistente.

```js
assert.lengthOf(document.querySelectorAll('article'),1);
```

Il tuo elemento `article` dovrebbe avere due elementi `p`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
assert.lengthOf(paragraphChildren,2);
```

Il testo del tuo primo elemento `p` dovrebbe essere `French Vanilla`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const firstP = paragraphChildren?.[0];
assert.match(firstP?.innerText.trim(), /French Vanilla/i);
```

Il testo del tuo secondo elemento `p` dovrebbe essere `3.00`.

```js
const article = document.querySelector('article');
const paragraphChildren = article?.querySelectorAll(`:scope ${'p'}`);
const secondP = paragraphChildren?.[1];
assert.match(secondP?.innerText.trim(), /3\.00/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
--fcc-editable-region--
            
--fcc-editable-region--
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
