---
id: 5f3c866d5414453fc2d7b480
title: Passaggio 27
challengeType: 0
dashedName: step-27
---

# --description--

Inizia sotto la coppia caffè/prezzo esistente, aggiungendo i seguenti caffè e prezzi usando elementi `article` con due elementi `p` annidati all'interno di ciascuno.

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

Come prima, il testo del primo elemento `p` dovrebbe contenere il gusto del caffè e il testo del secondo elemento `p` dovrebbe contenere il prezzo.

# --before-each--

```js
const articles = document.querySelectorAll('article');
```

# --hints--

Dovresti avere cinque elementi `article`.

```js
assert.lengthOf(articles, 5);
```

Ogni elemento `article` dovrebbe avere due elementi `p`.

```js
assert.isNotEmpty(articles);
articles.forEach(article => {
  const children = [...article.children];
  assert.lengthOf(children, 2);
  assert.isTrue(children.every(el => el.tagName === 'P'));
});
```

Il tuo primo elemento `article` dovrebbe avere elementi `p` con il testo `French Vanilla` e `3.00`.

```js
const children = articles?.[0]?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

Il tuo secondo elemento `article` dovrebbe avere elementi `p` con il testo `Caramel Macchiato` e `3.75`.

```js
const children = articles?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

Il tuo terzo elemento `article` dovrebbe avere elementi `p` con il testo `Pumpkin Spice` e `3.50`.

```js
const children = articles?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

Il tuo quarto elemento `article` dovrebbe avere elementi `p` con il testo `Hazelnut` e `4.00`.

```js
const children = articles?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

Il tuo quinto elemento `article` dovrebbe avere elementi `p` con il testo `Mocha` e `4.50`.

```js
const children = articles?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
