---
id: 5f3ef6e0a81099d9a697b550
title: Passaggio 61
challengeType: 0
dashedName: step-61
---

# --description--

All'interno del `address`, aggiungi un elemento `p`. Poi, annida un elemento àncora (`a`) nel `p` che colleghi a `https://www.freecodecamp.org` e abbia il testo `Visit our website`.

Assicurati che il collegamento si apra in una nuova scheda aggiungendo un attributo `target` con il valore `_blank`.

# --hints--

Non dovresti modificare l'elemento `footer` esistente.

```js
assert.lengthOf(document.querySelectorAll("footer"), 1);
```

Non dovresti modificare l'elemento `address` esistente.

```js
assert.lengthOf(document.querySelectorAll("address"), 1);
```

Il tuo nuovo elemento `p` dovrebbe essere annidato all'interno del tuo elemento `address`. Dovresti avere esattamente un elemento `p`.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p"), 1);
assert.lengthOf(document.querySelectorAll("footer p"), 1);
```

Il tuo nuovo elemento `a` dovrebbe essere annidato all'interno del tuo nuovo elemento `p`. Dovresti avere esattamente un elemento `a`.

```js
assert.lengthOf(document.querySelectorAll("footer > address > p > a"), 1);
assert.lengthOf(document.querySelectorAll("footer a"), 1);
```

Il tuo nuovo elemento `a` dovrebbe avere il testo `Visit our website`.

```js
assert.equal(document.querySelector("footer > address > p > a")?.innerText.trim(), "Visit our website");
```

Il tuo nuovo elemento `a` dovrebbe collegare a `https://www.freecodecamp.org`. Ricorda che gli elementi `a` usano l'attributo `href` per creare un collegamento.

```js
assert.equal(document.querySelector("footer > address > p > a")?.href, "https://www.freecodecamp.org/");
```

Il tuo nuovo elemento `a` dovrebbe avere l'attributo `target` impostato su `_blank`.

```js
assert.equal(document.querySelector("footer > address > p > a")?.target, "_blank");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p class="established">Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="flavor">Caramel Macchiato</p><p class="price">3.75</p>
          </article>
          <article class="item">
            <p class="flavor">Pumpkin Spice</p><p class="price">3.50</p>
          </article>
          <article class="item">
            <p class="flavor">Hazelnut</p><p class="price">4.00</p>
          </article>
          <article class="item">
            <p class="flavor">Mocha</p><p class="price">4.50</p>
          </article>
        </section>
        <section>
          <h2>Desserts</h2>
          <article class="item">
            <p class="dessert">Donut</p><p class="price">1.50</p>
          </article>
          <article class="item">
            <p class="dessert">Cherry Pie</p><p class="price">2.75</p>
          </article>
          <article class="item">
            <p class="dessert">Cheesecake</p><p class="price">3.00</p>
          </article>
          <article class="item">
            <p class="dessert">Cinnamon Roll</p><p class="price">2.50</p>
          </article>
        </section>
      </main>
      <footer>
        <address>
--fcc-editable-region--
          
--fcc-editable-region--
        </address>
      </footer>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
  font-family: sans-serif;
}

h1 {
  font-size: 40px;
}

h2 {
  font-size: 30px;
}

.established {
  font-style: italic;
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
  padding: 20px;
  max-width: 500px;
}

h1, h2 {
  font-family: Impact, serif;
}

.item p {
  display: inline-block;
}

.flavor, .dessert {
  text-align: left;
  width: 75%;
}

.price {
  text-align: right;
  width: 25%
}
```
