---
id: 66c6491fe4c8e0f16845425f
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

حان الوقت الآن لإضافة ثمرة أخرى إلى القائمة.

باستخدام نفس طريقة المصفوفة كما في السابق، أضف السلسلة النصية `"Grapes"` إلى نهاية مصفوفة `shoppingList`.

ثم أضف `console.log` واستدعي الدالة `getShoppingListMsg` مع مصفوفة `shoppingList` كمعلمة داخل `console.log` لرؤية القائمة المحدثة في وحدة التحكم.

# --hints--

يجب أن تحتوي مصفوفة `shoppingList` الآن على عنصرين.

```js
assert.lengthOf(shoppingList, 2);
```

يجب أن تحتوي مصفوفة `shoppingList` على السلسلة النصية `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

يجب استدعاء الدالة `getShoppingListMsg(shoppingList)` داخل `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
