---
id: 66c64095f61166ede6e39a84
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

في قائمة التسوق هذه، ستستخدم مصفوفة لتمثيل العناصر التي تحتاج إلى شرائها.

في الدروس السابقة، تعلمت كيفية إنشاء مصفوفات مثل هذه:

```js
const fruits = ['apple', 'banana', 'orange'];
```

في هذه الخطوة، أنشئ متغيرًا يسمى `shoppingList` وقم بتعيين مصفوفة فارغة له.

# --hints--

يجب أن يكون لديك متغير يسمى `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

يجب تعيين مصفوفة إلى المتغير `shoppingList`.

```js
assert.isArray(shoppingList);
```

يجب أن تكون المصفوفة الخاصة بك فارغة.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
