---
id: 66c63bfa4685e4e3843fa0dc
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستواصل تعلم المصفوفات من خلال بناء قائمة تسوق البقالة.

ابدأ بإضافة `console.log` يسجل السلسلة النصية `"Grocery shopping list"` في وحدة التحكم.

# --hints--

يجب أن يكون لديك `console.log` في شفرتك.

```js
assert.match(code, /console\.log(.*)/);
```

يجب أن يعرض `console.log` الرسالة `"Grocery shopping list"`. تحقق جيدًا من الأخطاء الإملائية وحالة الأحرف.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
