---
id: 66c645b345a39bf04864dc50
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

لعرض قائمة التسوق المحدثة، يجب تسجيل قائمة التسوق الحالية ورسالة قصيرة في وحدة التحكم.

بما أن هذه الرسالة ستُستخدم بشكل متكرر طوال ورشة العمل، من الأفضل إنشاء دالة قابلة لإعادة الاستخدام.

أنشئ دالة باسم `getShoppingListMsg` تأخذ مصفوفة كمعلمة وتُرجع السلسلة النصية `"Current Shopping List: "` متبوعة بمحتويات المصفوفة المقدمة.

يمكنك استخدام القوالب الحرفية أو دمج السلاسل النصية باستخدام عامل `+` لربط السلسلة بالمصفوفة الممررة إلى الدالة.

# --hints--

يجب أن تحتوي شفرتك على دالة باسم `getShoppingListMsg`.

```js
assert.isFunction(getShoppingListMsg);
```

يجب أن تُرجع دالة `getShoppingListMsg` سلسلة نصية.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

يجب أن تقبل دالة `getShoppingListMsg` مصفوفة وتُرجع الرسالة `"Current Shopping List: "` متبوعة بمحتويات تلك المصفوفة. تحقق جيدًا من المسافات والأخطاء الإملائية.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
