---
id: 66c644c08ebcd2ef90c750bd
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

في الدروس السابقة، تعلمت كيف تضيف عناصر إلى نهاية مصفوفة باستخدام الطريقة `push` مثل هذا:

```js
fruits.push('pear');
```

باستخدام الطريقة `push`، أضف السلسلة النصية `"Apples"` إلى مصفوفة `shoppingList`.

# --hints--

يجب أن تحتوي مصفوفة `shoppingList` على عنصر واحد.

```js
assert.lengthOf(shoppingList, 1);
```

يجب أن تحتوي مصفوفة `shoppingList` على السلسلة النصية `"Apples"`. تحقق جيدًا من الأخطاء الإملائية وأحرف الكتابة.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
