---
id: 66c64210413532ee9d3bd342
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في الجزء التالي من ورشة العمل، ستتدرب على إضافة عناصر طعام إلى قائمة البقالة.

ابدأ باستخدام `console.log` لتسجيل الرسالة `"It will be nice to have some fruit to eat."`.

# --hints--

يجب أن تسجل الرسالة `"It will be nice to have some fruit to eat."` في وحدة التحكم.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
