---
id: 66c72b0ffbc5522525768558
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

في الخطوات السابقة، راجعت كيفية إضافة عنصر إلى نهاية المصفوفة باستخدام طريقة `push`.

لكن طريقة `push` تقبل عدة معلمات، لذا يمكنك إضافة عدة عناصر إلى نهاية المصفوفة بهذه الطريقة:

```js
array.push(item1, item2, item3);
```

في هذه الخطوة، استخدم طريقة `push` لإضافة السلاسل النصية `"Popcorn", "Beef Jerky", "Potato Chips"` إلى مصفوفة `shoppingList`.

الترتيب مهم، فتأكد من إضافة العناصر بالترتيب الذي وردت به.

# --hints--

يجب أن تحتوي مصفوفة `shoppingList` على إجمالي `6` سلاسل نصية.

```js
assert.lengthOf(shoppingList, 6);
```

يجب أن تحتوي مصفوفة `shoppingList` على السلسلة النصية `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

يجب أن تحتوي مصفوفة `shoppingList` على السلسلة النصية `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

يجب أن تحتوي مصفوفة `shoppingList` على السلسلة النصية `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
