---
id: 66c730ee6ae076281721d0b9
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

في الدروس السابقة، تعلمت كيف تزيل عناصر من نهاية المصفوفة باستخدام الطريقة `pop`.

إليك تذكير بكيفية استخدام الطريقة `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

استخدم الطريقة `pop` لإزالة العنصر الأخير من المصفوفة `shoppingList`.

# --hints--

يجب أن تحتوي مصفوفة `shoppingList` الآن على `5` عناصر.

```js
assert.lengthOf(shoppingList, 5);
```

يجب ألا تحتوي المصفوفة `shoppingList` على السلسلة النصية `"Potato Chips"` بعد الآن.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
