---
id: 66c73a0c5b264f2a75164d94
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

حان الوقت الآن لتسجيل مصفوفة `shoppingList` المحدثة في وحدة التحكم.

أضف `console.log` واستدعِ الدالة `getShoppingListMsg` مع مصفوفة `shoppingList` كمعلمة داخل `console.log` لرؤية القائمة المحدثة مسجلة في وحدة التحكم.

# --hints--

يجب أن تستدعي `getShoppingListMsg(shoppingList)` داخل `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 5);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();

--fcc-editable-region--

--fcc-editable-region--
```
