---
id: 66c726c34ecf1e238aa9d7d5
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

في الدروس السابقة، تعلمت كيفية إضافة عناصر إلى بداية مصفوفة باستخدام الطريقة `unshift()`.

إليك تذكير بكيفية استخدام الطريقة `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

استخدم الطريقة `unshift()` لإضافة السلسلة النصية `"Vegetable Oil"` إلى بداية مصفوفة `shoppingList`.

# --hints--

يجب أن تحتوي مصفوفة `shoppingList` على ثلاثة عناصر.

```js
assert.lengthOf(shoppingList, 3);
```

يجب أن تستخدم الطريقة `unshift()` لإضافة `"Vegetable Oil"` إلى بداية مصفوفة `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
