---
id: 66c73a7798f6f62b2ae58f22
title: الخطوة 16
challengeType: 1
dashedName: step-16
---

# --description--

حان الوقت الآن لإضافة المزيد من العناصر إلى بداية قائمة البقالة.

ابدأ بإضافة عبارة `console.log` تسجل الرسالة `"It might be nice to get a dessert."`

أسفل عبارة `console` تلك، استخدم الطريقة الصحيحة للمصفوفة لإضافة السلسلة النصية `"Chocolate Cake"` إلى بداية مصفوفة `shoppingList`.

أخيرًا، أضف `console.log` واستدعِ الدالة `getShoppingListMsg` مع مصفوفة `shoppingList` كمعلمة داخل `console.log` لرؤية القائمة المحدثة مسجلة في وحدة التحكم.

# --hints--

يجب أن يكون لديك `console.log` يسجل الرسالة `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

يجب أن تحتوي مصفوفة `shoppingList` على إجمالي `6` عناصر.

```js
assert.lengthOf(shoppingList, 6);
```

يجب أن تستخدم الطريقة `unshift` لإضافة السلسلة النصية `"Chocolate Cake"` إلى بداية مصفوفة `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

يجب أن تستدعي `getShoppingListMsg(shoppingList)` داخل `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
