---
id: 66c742d045c9fc2e09fa64b1
title: الخطوة 19
challengeType: 1
dashedName: step-19
---

# --description--

آخر تعديل يجب إجراؤه على قائمة البقالة هو تحديث العنصر الأول في القائمة.

في الدروس السابقة، تعلمت كيفية تحديث عنصر باستخدام ترميز الأقواس والفهرس الخاص بالعنصر الذي تريد تحديثه.

إليك تذكير بكيفية تحديث عنصر في مصفوفة:

```js
const array = [1, 2, 3, 4, 5];

array[0] = 10;
// Result: [10, 2, 3, 4, 5]
console.log(array); 
```

حدّث العنصر الأول في مصفوفة `shoppingList` ليصبح `"Canola Oil"`.

# --hints--

يجب أن تحدّث العنصر الأول في مصفوفة `shoppingList` ليصبح `"Canola Oil"`.

```js
assert.equal(shoppingList[0], "Canola Oil");
assert.notEqual(shoppingList[1], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

console.log("It might be nice to get a dessert.");

shoppingList.unshift("Chocolate Cake");
console.log(getShoppingListMsg(shoppingList));

console.log("On second thought, maybe we should be more health conscious.");

shoppingList.shift();

--fcc-editable-region--

--fcc-editable-region--
```
