---
id: 66cbe2319d3845545a293a0b
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

حان الوقت الآن لرؤية الرسالة المسجلة في وحدة التحكم.

أضف `console.log` واستدعي الدالة `getShoppingListMsg` مع المصفوفة `shoppingList` كمعلمة داخل `console.log` لرؤية الرسالة المسجلة في وحدة التحكم.

# --hints--

يجب استدعاء الدالة `getShoppingListMsg(shoppingList)` داخل `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
