---
id: 66f686b8ebdb982fa8e14330
title: Створіть програму для конвертації дати
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

У цій лабораторній роботі ви напишете програму для конвертації дати, яка перетворює задану дату у різні формати. Наприклад, поточна дата `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` буде відформатована у такі 2 способи:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історія користувача:**

1. Ви повинні створити змінну з ім’ям `currentDate` і призначити їй поточну дату та час, використовуючи об’єкт `Date`.
2. Ви повинні створити змінну з ім’ям `currentDateFormat` і призначити їй рядок `Current Date and Time: [current date]`. Замініть `[current date]` на результат змінної `currentDate`.
3. Ви повинні вивести значення `currentDateFormat` у консоль.
4. Ви повинні створити функцію з ім’ям `formatDateMMDDYYYY`, яка приймає об’єкт дати як параметр. Ви можете назвати цей параметр як завгодно.
5. Ваша функція `formatDateMMDDYYYY` повинна конвертувати поточну дату у формат `month/day/year` і повертати рядок `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Ви повинні створити функцію з ім’ям `formatDateLong`, яка приймає об’єкт дати як параметр.
7. Ваша функція `formatDateLong` повинна конвертувати поточну дату у формат `Month Day, Year` і повертати рядок `Formatted Date (Month Day, Year): [formatted date]`.

**Примітка:** Щоб тести пройшли успішно, переконайтеся, що ви використовуєте `en-US` як локаль при форматуванні дат.

# --hints--

Ви повинні мати змінну з ім’ям `currentDate`, яка містить поточну дату та час, використовуючи об’єкт `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Ви повинні мати змінну з ім’ям `currentDateFormat`, яка містить поточну дату у форматі `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Ви повинні вивести значення `currentDateFormat` у консоль.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Ви повинні мати функцію `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

Функція `formatDateMMDDYYYY` повинна приймати один параметр.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Коли об’єкт дати містить `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, функція `formatDateMMDDYYYY` повинна повертати `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Ви повинні мати функцію `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

Функція `formatDateLong` повинна приймати один параметр.

```js
assert.lengthOf(formatDateLong, 1);
```

Коли об’єкт дати містить `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, функція `formatDateLong` повинна повертати `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
