---
id: 69aaafeaa4037197c75587cd
title: Реалізувати чергу
challengeType: 26
dashedName: implement-a-queue
---

# --description--

У цій лабораторній роботі ви реалізуєте структуру даних черга за допомогою функцій. Черга — це структура даних типу «перший прийшов — перший пішов» (FIFO), де елементи додаються в кінець черги і видаляються з початку.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати функцію `enqueue`, яка додає елемент у кінець черги.
1. Ви повинні мати функцію `dequeue`, яка видаляє і повертає перший елемент черги.
1. Ви повинні мати функцію `front`, яка повертає перший елемент черги без видалення.
1. Ви повинні мати функцію `size`, яка повертає кількість елементів у черзі.
1. Ви повинні мати функцію `isEmpty`, яка повертає `true`, якщо черга порожня, і `false` в іншому випадку.

# --hints--

Ви повинні мати функцію `enqueue`.

```js
assert.isFunction(enqueue);
```

Функція `enqueue` повинна додавати елемент у кінець черги.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

Ви повинні мати функцію `dequeue`.

```js
assert.isFunction(dequeue);
```

Функція `dequeue` повинна видаляти і повертати перший елемент черги.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(dequeue(queue), 'first');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(size(queue), 1);
```

`dequeue` має повертати `undefined`, коли викликається на порожній черзі.

```js
const queue = initQueue();
assert.isUndefined(dequeue(queue));
```

`enqueue` має продовжувати додавати елементи в кінець черги після видалень.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(dequeue(queue), 'first');
enqueue(queue, 'third');
assert.strictEqual(front(queue), 'second');
assert.strictEqual(dequeue(queue), 'second');
assert.strictEqual(dequeue(queue), 'third');
```

Різні екземпляри черги не повинні ділити стан.

```js
const queueOne = initQueue();
const queueTwo = initQueue();
enqueue(queueOne, 'first');
enqueue(queueTwo, 'second');
assert.strictEqual(front(queueOne), 'first');
assert.strictEqual(front(queueTwo), 'second');
assert.strictEqual(size(queueOne), 1);
assert.strictEqual(size(queueTwo), 1);
```

Ви повинні мати функцію `front`.

```js
assert.isFunction(front);
```

Функція `front` повинна повертати перший елемент черги без видалення.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(front(queue), 'first');
assert.strictEqual(size(queue), 2);
```

`front` має повертати `undefined`, коли викликається на порожній черзі.

```js
const queue = initQueue();
assert.isUndefined(front(queue));
```

Ви повинні мати функцію `size`.

```js
assert.isFunction(size);
```

Функція `size` повинна повертати правильну кількість елементів у черзі.

```js
const queue = initQueue();
assert.strictEqual(size(queue), 0);
enqueue(queue, 'first');
assert.strictEqual(size(queue), 1);
enqueue(queue, 'second');
enqueue(queue, 'third');
assert.strictEqual(size(queue), 3);
dequeue(queue);
assert.strictEqual(size(queue), 2);
```

Ви повинні мати функцію `isEmpty`.

```js
assert.isFunction(isEmpty);
```

Функція `isEmpty` повинна повертати `true` для порожньої черги.

```js
const queue = initQueue();
assert.isTrue(isEmpty(queue));
```

Функція `isEmpty` повинна повертати `false` для непорожньої черги.

```js
const queue = initQueue();
enqueue(queue, 'element');
assert.isFalse(isEmpty(queue));
```

Функція `isEmpty` повинна повертати `true` після видалення всіх елементів із черги.

```js
const queue = initQueue();
enqueue(queue, 'first');
enqueue(queue, 'second');
dequeue(queue);
dequeue(queue);
assert.isTrue(isEmpty(queue));
```

# --seed--

## --seed-contents--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  
}

function dequeue(queue) {
  
}

function front(queue) {
  
}

function size(queue) {
  
}

function isEmpty(queue) {
  
}
```

# --solutions--

```js
function initQueue() {
  return {
    collection: []
  };
}

function print(queue) {
  console.log(queue.collection);
}

function enqueue(queue, element) {
  queue.collection.push(element);
}

function dequeue(queue) {
  return queue.collection.shift();
}

function front(queue) {
  return queue.collection[0];
}

function size(queue) {
  return queue.collection.length;
}

function isEmpty(queue) {
  return queue.collection.length === 0;
}
```
