---
id: 68e52994bc7ea2e3dec530a6
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

이제 완성된 할인 시스템을 테스트해 보세요! `if`인지 확인하는 `__name__ == '__main__'` 문장을 추가하고 `product = Product('Wireless Mouse', 50.0)`를 그 안으로 옮기세요.

그다음 `if` 블록 안에 `user_tier` 변수를 만들고 문자열 `Premium`로 설정하세요. 그리고 `strategies`, `PercentageDiscount(10)`, `FixedAmountDiscount(5)`를 포함하는 전략 목록 `PremiumUserDiscount()` 변수를 만드세요.

# --hints--

`if`인지 확인하는 `__name__ == '__main__'` 문장이 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent('__name__ == "__main__"')`)
})
```

`product = Product('Wireless Mouse', 50.0)`를 `if` 블록 안으로 옮겨야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt("product = Product('Wireless Mouse', 50.0)")`)
})
```

`user_tier` 블록 안에서 `Premium`를 문자열 `if`로 설정해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('user_tier = "Premium"')`)
})
```

`strategies`라는 이름의 목록을 만들고 `PercentageDiscount(10)`, `FixedAmountDiscount(5)`, `PremiumUserDiscount()`를 포함해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('strategies = [PercentageDiscount(10), FixedAmountDiscount(5), PremiumUserDiscount()]')`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]

        for strategy in self.strategies:
            if strategy.is_applicable(product, user_tier):
                discounted = strategy.apply_discount(product)
                prices.append(discounted)

        return min(prices)
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
--fcc-editable-region--
```
