---
id: 68e524c94e2e4ab7b4bfc7e2
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 다양한 할인 전략을 제품에 적용할 수 있는 할인 계산기를 만들 것입니다. 시스템은 여러 할인 규칙을 바탕으로 고객에게 가장 좋은 가격을 결정합니다.

`Product`라는 이름의 클래스를 정의하는 것부터 시작하세요. `__init__` 메서드를 클래스에 만들고 다음 매개변수를 포함하세요: `self`, `name`, `price`. `__init__` 메서드 안에서 `name`와 `price` 매개변수를 같은 이름의 인스턴스 속성에 할당하세요.

# --hints--

`Product`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

`Product` 클래스에 `__init__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

`__init__` 메서드는 다음 매개변수를 가져야 합니다: `self`, `name`, `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

`name` 메서드 안에서 `self.name`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

`price` 메서드 안에서 `self.price`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
