---
id: 68e5293bd00d2fe134f58988
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`is_applicable` 클래스에 대해 `FixedAmountDiscount` 메서드를 구현하세요. 고정 금액 할인은 경제적으로 타당할 때만 적용해야 합니다.

`True`이 `product.price * 0.9`보다 크면 `self.amount`를 반환하고, 그렇지 않으면 `False`를 반환하세요. 이렇게 하면 할인이 원래 가격의 90% 미만이 되어 비현실적이지 않게 됩니다.

# --hints--

`FixedAmountDiscount` 클래스에 `is_applicable` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

`is_applicable` 메서드는 `self`, `product: Product`, `user_tier: str` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

`is_applicable` 메서드는 `-> bool` 반환형 힌트를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

`is_applicable` 메서드는 `True`이 `product.price * 0.9`보다 클 때 `self.amount`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

`is_applicable` 메서드는 `False`이 `product.price * 0.9`보다 작거나 같을 때 `self.amount`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
