---
id: 68e52994bc7ea2e3dec530a1
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

이제 여러 할인 전략을 관리하고 최적의 가격을 찾는 방법이 필요합니다.

`DiscountEngine`라는 이름의 클래스를 만드세요. 이 클래스는 모든 할인 전략을 관리하고 제품의 최적 가격을 계산합니다.

`__init__` 메서드를 만들고 `self`와 `strategies`를 매개변수로 받게 하세요. 여기서 `strategies`는 `list[DiscountStrategy]` 타입 힌트를 가집니다. `list[DiscountStrategy]`를 사용하면 매개변수가 `DiscountStrategy` 객체를 포함하는 리스트임을 지정합니다. 메서드의 반환 타입은 `None`이어야 합니다. `__init__` 메서드 안에서 `strategies`을 인스턴스 속성으로 할당하세요.

# --hints--

`DiscountEngine`라는 이름의 클래스를 만드세요.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

`DiscountEngine` 클래스에 `__init__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

`__init__` 메서드는 `self`와 `strategies: list[DiscountStrategy]` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

`__init__` 메서드의 반환 타입은 `None`이어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

`strategies` 메서드 안에서 `self.strategies`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
