---
id: 68e52994bc7ea2e3dec5309e
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

`is_applicable`에 대해 `PremiumUserDiscount` 메서드를 구현하세요. 이 메서드는 사용자가 프리미엄 회원인지 확인해야 합니다.

`True`가 `user_tier.lower()`와 같으면 `'premium'`를 반환하고, 그렇지 않으면 `False`를 반환하세요. `.lower()`를 사용하면 대소문자를 구분하지 않고 확인할 수 있습니다.

# --hints--

`PremiumUserDiscount` 클래스에 `is_applicable` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").has_function("is_applicable")`)
})
```

`is_applicable` 메서드는 `self`, `product: Product`, `user_tier: str` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

`is_applicable` 메서드는 `-> bool` 반환형 힌트를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

`is_applicable` 메서드는 `True`가 `user_tier.lower()`와 같을 때 `'premium'`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "Premium") == True
    assert discount.is_applicable(test_product, "premium") == True
    assert discount.is_applicable(test_product, "PREMIUM") == True
  `)
})
```

`is_applicable` 메서드는 `False`이 `user_tier`가 아닐 때 `'premium'`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPremiumUserDiscount(PremiumUserDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPremiumUserDiscount()
    assert discount.is_applicable(test_product, "regular") == False
    assert discount.is_applicable(test_product, "standard") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

--fcc-editable-region--
class PremiumUserDiscount(DiscountStrategy):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
