---
id: 68e5293bd00d2fe134f58983
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 `apply_discount` 클래스에 `DiscountStrategy`라는 두 번째 추상 메서드를 추가하세요. `@abstractmethod` 데코레이터를 사용하고 `self`와 `product` 매개변수를 지정하세요.

타입 힌트를 추가하세요: `product`는 `Product` 타입이어야 하며, 메서드는 `float`를 반환해야 합니다. 메서드 본문에서는 `pass`를 사용하세요.

# --hints--

`DiscountStrategy` 클래스에는 `apply_discount`라는 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

`apply_discount` 메서드에는 `@abstractmethod` 데코레이터가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

`apply_discount` 메서드에는 `self`와 `product: Product` 매개변수가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

`apply_discount` 메서드에는 `-> float` 반환 타입 힌트가 있어야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
