---
id: 68e529226b0defdff26c4a8b
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이제 할인 시스템 인터페이스를 구축하기 시작할 거예요. Python은 다른 클래스의 청사진 역할을 하는 추상 클래스를 정의하는 방법을 제공합니다. 이는 `ABC` 모듈의 `abc` 클래스(Abstract Base Class)를 사용해 구현합니다.

코드 상단에 다음 가져오기 문을 추가하세요:

```py
from abc import ABC, abstractmethod
```

이 코드는 `ABC` 클래스와 `abstractmethod` 데코레이터를 가져오는데, 추상 클래스를 상속하는 모든 클래스가 반드시 구현해야 하는 메서드를 정의할 때 사용할 거예요.

# --hints--

코드 상단에 언급한 가져오기 문이 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
