---
id: 68e52994bc7ea2e3dec5309f
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

이제 `apply_discount`에 대해 `PremiumUserDiscount` 메서드를 구현하세요. 이 메서드는 프리미엄 사용자에게 20% 할인을 적용해야 합니다.

원래 가격의 80%, 즉 20% 할인된 가격을 나타내는 `product.price * 0.8`를 반환하세요.

# --hints--

`PremiumUserDiscount` 클래스에는 `apply_discount` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").has_function("apply_discount")`)
})
```

`apply_discount` 메서드는 `self`와 `product: Product` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

`apply_discount` 메서드는 `-> float` 반환 타입 힌트를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").find_function("apply_discount").has_returns("float")`)
})
```

`apply_discount` 메서드는 `product.price * 0.8`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PremiumUserDiscount()
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

`apply_discount` 메서드는 20% 할인(원래 가격의 80%)을 적용해야 합니다.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 50.0)
    discount = PremiumUserDiscount()
    result = discount.apply_discount(test_product)
    assert result == 40.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
