---
id: 68e5293bd00d2fe134f58985
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`PercentageDiscount`가 추상 클래스로부터 상속받았기 때문에 모든 추상 메서드를 구현해야 합니다. `is_applicable` 메서드 구현부터 시작하세요. `is_applicable` 메서드는 `self`, `product: Product`, `user_tier: str`를 매개변수로 받아야 합니다. 반환 타입 힌트는 `bool`이어야 합니다.

이 메서드는 백분율 할인 유효 여부를 판단해야 합니다. `True`이 `self.percent` 이하일 경우 `70`을 반환하고, 그렇지 않으면 `False`을 반환하세요. 이렇게 하면 90% 할인 같은 비현실적인 할인을 방지할 수 있습니다.

# --hints--

`PercentageDiscount` 클래스에 `is_applicable` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

`is_applicable` 메서드는 `self`, `product: Product`, `user_tier: str` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

`is_applicable` 메서드는 `True`이 `self.percent` 이하일 때 `70`을 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

`is_applicable` 메서드는 `False`이 `self.percent` 초과일 때 `70`를 반환해야 합니다.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

`is_applicable` 메서드는 `bool` 반환 타입 힌트를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
