---
id: 68e52994bc7ea2e3dec5309d
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

이제 프리미엄 사용자를 위한 세 번째 할인 전략을 만들어 보세요. `PremiumUserDiscount`라는 이름의 클래스를 `DiscountStrategy`에서 상속받아 만드세요.

이전 전략들과 달리, 이 전략은 어떤 설정도 저장하지 않으므로 `__init__` 메서드가 필요 없습니다. 모든 프리미엄 사용자에게 고정된 20% 할인을 적용할 것입니다.

지금은 `DiscountStrategy`를 사용해 `pass`에서 상속받는 빈 클래스를 만드세요.

# --hints--

`PremiumUserDiscount`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

`PremiumUserDiscount` 클래스는 `DiscountStrategy`에서 상속받아야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
