---
id: 68e52994bc7ea2e3dec530a4
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

이제 `for` 반복문을 만들어 `self.strategies`을 반복하세요. 각 `strategy`에 대해 `is_applicable` 메서드를 사용해 적용 가능한지 확인하세요. 적용 가능하면 `apply_discount`를 사용해 할인된 가격을 계산하고 `discounted`라는 변수에 저장하세요. 그런 다음 할인된 가격을 `prices` 목록에 추가하세요.

# --hints--

`for` 반복문을 만들어 `self.strategies`을 반복해야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_for_iter().is_equivalent("self.strategies")
  `)
})
```

반복문 안에서 `if` 문장은 `if strategy.is_applicable(product, user_tier):`이어야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_conditions()[0].is_equivalent("strategy.is_applicable(product, user_tier)")
  `)
})
```

전략이 적용 가능하면 `apply_discount`를 호출하고 결과를 `prices`에 추가하세요.

```js
({
  test: () => runPython(`
assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").find_for_loops()[0].find_ifs()[0].find_body().is_ordered("discounted = strategy.apply_discount(product)","prices.append(discounted)")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

--fcc-editable-region--
    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]
        

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
