---
id: 68e5293bd00d2fe134f5898a
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

새 할인 전략을 테스트해보세요. 기존 코드 뒤에 `FixedAmountDiscount`의 인스턴스를 `5` 금액으로 생성하고 `fixed_discount`라는 변수에 할당하세요. 그런 다음 이 할인을 제품에 적용하고 결과를 출력하세요.

# --hints--

`fixed_discount`라는 변수에 `FixedAmountDiscount` 인스턴스를 생성해야 합니다.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("fixed_discount")
    assert isinstance(fixed_discount, FixedAmountDiscount)
  `)
})
```

`FixedAmountDiscount` 인스턴스는 `5` 금액으로 초기화해야 합니다.

```js
({
  test: () => runPython(`assert fixed_discount.amount == 5`)
})
```

`fixed_discount.apply_discount(product)`을 호출하고 결과를 출력해야 합니다.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_call("print(fixed_discount.apply_discount(product))")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
--fcc-editable-region--

--fcc-editable-region--
```
