---
id: 68e5293bd00d2fe134f58982
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이전 수업에서 설명했듯이, 추상 메서드는 추상 기본 클래스를 상속받는 모든 클래스가 반드시 구현해야 하는 메서드입니다. 이들은 모든 할인 전략이 특정 필수 메서드를 갖도록 보장하는 계약 역할을 합니다.

`DiscountStrategy` 클래스 내부에 `is_applicable`라는 이름의 메서드를 정의하세요. 메서드 정의 위에 `@abstractmethod` 데코레이터를 사용하세요. 메서드는 `self`, `product`, `user_tier` 매개변수를 가져야 합니다. 메서드 본문에는 `pass` 문장을 사용하세요.

메서드 시그니처를 명확히 하기 위해 타입 힌트를 추가하세요: `product`는 `Product` 타입이어야 하고, `user_tier`는 `str` 타입이어야 합니다.


# --hints--

`DiscountStrategy` 클래스에는 `is_applicable`라는 이름의 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

`is_applicable` 메서드에는 `@abstractmethod` 데코레이터가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

`is_applicable` 메서드에는 `self`, `product: Product`, `user_tier: str` 매개변수가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
