---
id: 68e5293bd00d2fe134f58984
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 추상 클래스를 구체적으로 구현할 차례입니다. `PercentageDiscount`라는 이름의 클래스를 만들고 `DiscountStrategy`에서 상속받으세요.

`__init__`와 `self`를 매개변수로 받는 `percent` 메서드를 만드세요. `int`에는 `percent` 타입 힌트를 주고, 메서드의 반환 타입은 `None`로 설정하세요. `__init__` 메서드 안에서 `percent`를 인스턴스 속성으로 저장하세요.

# --hints--

`PercentageDiscount`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

`PercentageDiscount` 클래스는 `DiscountStrategy`에서 상속받아야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

`PercentageDiscount` 클래스에는 `__init__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

`__init__` 메서드는 `self`와 `percent: int` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

`__init__` 메서드는 `None` 반환 타입 힌트를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

`percent` 메서드 안에서 `self.percent`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
