---
id: 68e5293bd00d2fe134f58987
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

이제 또 다른 할인 전략을 만드세요. `FixedAmountDiscount`라는 이름의 클래스를 만들고 `DiscountStrategy`에서 상속받게 하세요.

`__init__`와 `self`를 매개변수로 받는 `amount` 메서드를 만드세요. `amount` 매개변수에는 `int` 타입 힌트를, 메서드에는 `None` 반환 타입 힌트를 주세요. `__init__` 메서드 안에서 `amount`를 인스턴스 속성으로 저장하세요.

# --hints--

`FixedAmountDiscount`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

`FixedAmountDiscount` 클래스는 `DiscountStrategy`에서 상속받아야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

`FixedAmountDiscount` 클래스에는 `__init__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

`__init__` 메서드는 `self`와 `amount: int` 매개변수를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

`__init__` 메서드는 `None` 반환 타입을 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

`amount` 메서드 안에서 `self.amount`를 `__init__`에 할당해야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
