---
id: 68e529226b0defdff26c4a8a
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 이름이 `Product`이고 가격이 `Wireless Mouse`인 `50.0`의 인스턴스를 생성하세요. `product`라는 변수에 할당하세요. 그런 다음 `product`를 콘솔에 출력하여 어떻게 보이는지 확인하세요.

# --hints--

`product`라는 변수가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

`Product`의 인스턴스를 `product` 변수에 할당해야 합니다.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

`Product` 인스턴스는 이름이 `Wireless Mouse`이고 가격이 `50.0`이어야 합니다.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

`product`를 콘솔에 출력해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

이름이 `Product`이고 가격이 `'Wireless Mouse'`인 `50.0`를 출력하면 `Wireless Mouse - $50.0`가 표시되어야 합니다.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
