---
id: 68e5293bd00d2fe134f58986
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 `apply_discount` 메서드를 `PercentageDiscount`에 구현하세요.

이 메서드는 할인된 가격을 계산하여 반환해야 합니다.

공식은 다음과 같습니다: `product.price * (1 - self.percent / 100)`

# --hints--

`PercentageDiscount` 클래스에 `apply_discount` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

`apply_discount` 메서드는 매개변수 `self`와 `product: Product`를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

`apply_discount` 메서드는 반환 타입 힌트로 `float`를 가져야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

`apply_discount` 메서드는 공식 `product.price * (1 - self.percent / 100)`을 사용해 할인된 가격을 계산해야 합니다.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
