---
id: 68e5293bd00d2fe134f58981
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`DiscountStrategy`에서 상속받는 `ABC`라는 이름의 클래스를 만드세요. 이 클래스는 모든 할인 전략의 인터페이스를 정의하는 추상 기본 클래스가 될 것입니다.

지금은 클래스 본문에 `pass`만 사용하여 빈 클래스를 만드세요.

# --hints--

`DiscountStrategy`라는 이름의 클래스를 만들어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

`DiscountStrategy` 클래스는 `ABC`에서 상속받아야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
