---
id: 68e529226b0defdff26c4a89
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

지금은 `Product`의 인스턴스를 만들고 출력하면, 별로 유용하지 않은 기본 객체 표현이 나타납니다. 예를 들어, `<__main__.Product object at 0x7f9b8c2d1d30>`처럼 보일 수 있습니다.

기본 객체 표현을 읽기 쉬운 문자열로 바꾸려면, `__str__` 클래스에 `Product` 메서드를 만들어 `name - $price` 형식의 문자열을 반환하세요 (여기서 `name`와 `price`는 해당 속성 값으로 교체되어야 합니다).

# --hints--

`Product` 클래스에 `__str__` 메서드가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

`__str__` 메서드는 `name - $price` 형식의 문자열을 반환해야 합니다.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
