---
id: 68e79d91daf8fb4808f93c47
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

`apply_discount` 메서드를 구현한 후, 할인 전략을 테스트하세요. `PercentageDiscount` 퍼센트로 `10`의 인스턴스를 생성하고 `discount`라는 변수에 할당하세요. 그런 다음 `discount.apply_discount(product)`를 사용해 할인을 적용하고, 결과를 출력하여 할인이 제대로 작동하는지 확인하세요.

# --hints--

`discount`라는 변수에 `PercentageDiscount` 인스턴스를 생성해야 합니다.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

`PercentageDiscount` 인스턴스는 `10` 퍼센트로 초기화해야 합니다.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

`discount.apply_discount(product)`를 호출하고 결과를 출력해야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
