---
id: 68e9207ddcba053c5cb24413
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

Python에서 타입 힌트는 변수나 함수의 데이터 타입이 무엇일지 다른 개발자에게 알려주는 선택적 신호입니다. 다음은 `str` 매개변수에 `name` 타입 힌트를 추가한 예시입니다.

```py
def greet(name: str):
    return 'Hello, ' + name
```

`__init__` 메서드 매개변수에 대해 `name` 매개변수는 `str` 타입 힌트를, `price` 매개변수는 `float` 타입 힌트를 갖도록 업데이트하세요.

# --hints--

`str` 클래스의 `name` 메서드에서 `float` 매개변수에 `price` 타입 힌트를, `__init__` 매개변수에 `Product` 타입 힌트를 추가하세요.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
