---
id: 68ef53ef276acfa40c83cb12
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`bool` 메서드에 `is_applicable` 반환 타입 힌트를 추가하여 `bool`를 반환한다는 것을 명확히 나타내세요.

# --hints--

`is_applicable` 메서드에는 `-> bool` 반환 타입 힌트가 있어야 합니다.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_returns("bool")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str):
        pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
