---
id: 68f119ad905b29ff59d858c8
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이전 단계와 마찬가지로, `__str__` 메서드에 이 메서드가 객체의 문자열 표현을 반환하므로 `str` 반환 타입 힌트를 추가하세요.

# --hints--

`__str__` 메서드에는 `str` 반환 타입 힌트가 있어야 합니다.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
