---
id: 69a92f0b9cc04bb0d5327bb5
title: 스택 구현하기
challengeType: 26
dashedName: implement-a-stack
---

# --description--

이 실습에서는 함수를 사용해 스택 자료 구조를 구현할 것입니다. 스택은 후입선출(LIFO) 자료 구조로, 요소가 위쪽에서 추가되고 제거됩니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과해 실습을 완료하세요.

**사용자 스토리:**

1. 스택의 위에 요소를 추가하는 `push` 함수를 구현해야 합니다.
1. 스택의 위 요소를 제거하고 반환하는 `pop` 함수를 구현해야 하며, 요소가 없으면 `undefined`를 반환해야 합니다.
1. 스택의 위 요소를 제거하지 않고 반환하는 `peek` 함수를 구현해야 하며, 요소가 없으면 `undefined`를 반환해야 합니다.
1. 스택에 요소가 없으면 `isEmpty`를 반환하고, 그렇지 않으면 `true`를 반환하는 `false` 함수를 구현해야 합니다.
1. 스택의 모든 요소를 제거하는 `clear` 함수를 구현해야 합니다.

**참고:** 대부분의 테스트는 `push` 함수에 의존합니다. `push`을 먼저 구현하세요. `pop`, `peek`, `isEmpty`, `clear` 테스트는 스택에 요소를 추가해야 하므로 먼저 구현해야 합니다.

# --hints--

`push` 함수를 구현해야 합니다.

```js
assert.isFunction(push);
```

`push` 함수는 스택의 위에 요소를 추가해야 합니다.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

`pop` 함수를 구현해야 합니다.

```js
assert.isFunction(pop);
```

`pop` 함수는 스택의 위 요소를 제거하고 반환해야 합니다.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

`pop` 함수는 스택이 비어 있으면 `undefined`를 반환해야 합니다.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` 함수는 거짓 같은 값을 올바르게 반환해야 합니다.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

`peek` 함수를 구현해야 합니다.

```js
assert.isFunction(peek);
```

`peek` 함수는 스택의 위 요소를 제거하지 않고 반환해야 합니다.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

`peek` 함수는 스택이 비어 있으면 `undefined`를 반환해야 합니다.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` 함수는 거짓 같은 값을 올바르게 반환해야 합니다.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

`isEmpty` 함수를 구현해야 합니다.

```js
assert.isFunction(isEmpty);
```

`isEmpty` 함수는 빈 스택에 대해 `true`를 반환해야 합니다.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` 함수는 비어 있지 않은 스택에 대해 `false`를 반환해야 합니다.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

`isEmpty` 함수는 위 요소가 거짓 같은 값일 때 `false`를 반환해야 합니다.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

`clear` 함수를 구현해야 합니다.

```js
assert.isFunction(clear);
```

`clear` 함수는 스택의 모든 요소를 제거해야 합니다.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
