---
id: 68e41084bb275539f72de9fe
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتدرب على العمل مع الطرائق `replace()`, `replaceAll()` و `repeat()` من خلال بناء تطبيق لتحويل السلاسل النصية.

ابدأ بإنشاء متغير اسمه `originalString` وقم بتعيين القيمة `"I love cats."` له.

# --hints--

يجب أن يكون لديك متغير اسمه `originalString`.

```js
assert.exists(originalString);
```

يجب أن يكون المتغير `originalString` سلسلة نصية.

```js
assert.isString(originalString);
```

يجب تعيين السلسلة النصية `"I love cats."` إلى المتغير `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
