---
id: 68e44519c769bf545bb3f22a
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

للحصول على الناتج المطلوب `"I love dogs and dogs are so much fun!"`، يجب استخدام الطريقة `replaceAll()` بدلًا من الطريقة `replace()`. تعيد هذه الطريقة سلسلة نصية جديدة لجميع التطابقات مع الجزء الفرعي.

حدّث `exampleSentence.replace("cats", "dogs")` لاستخدام الطريقة `replaceAll()` بدلًا من ذلك. سترى الآن النص الصحيح في وحدة التحكم.

# --hints--

يجب تحديث سطر الشفرة `exampleSentence.replace("cats", "dogs")` لاستخدام الطريقة `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
