---
id: 68e43f94365a964f26c88f78
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

حان الوقت الآن للعمل مع طريقة جديدة للسلسلة النصية.

ابدأ بإنشاء متغير جديد اسمه `exampleSentence` وقم بتعيين السلسلة النصية `"I love cats and cats are so much fun!"` له.

ثم أسفل ذلك المتغير، سجّل في وحدة التحكم `"Original sentence:"`. ثم أسفل ذلك الاستدعاء لـ `console.log()`، أضف استدعاءً آخر لـ `console.log()` مع المتغير `exampleSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير `exampleSentence`.

```js
assert.exists(exampleSentence);
```

يجب أن يكون المتغير `exampleSentence` سلسلة نصية.

```js
assert.isString(exampleSentence);
```

يجب أن تعيّن السلسلة النصية `"I love cats and cats are so much fun!"` إلى المتغير `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

يجب أن تسجل `"Original sentence:"` في وحدة التحكم.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

يجب أن تسجل المتغير `exampleSentence` في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
