---
id: 68e4465fba0bd256868b7deb
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

حان الوقت الآن للعمل مع طريقة جديدة للسلسلة النصية.

ابدأ بإنشاء متغير اسمه `learningSentence` وقم بتعيين السلسلة النصية `"I love learning!"` له.

ثم أسفل ذلك المتغير، سجّل في وحدة التحكم `"Original learning sentence:"`. أسفل ذلك الاستدعاء لـ `console.log()`، أضف استدعاءً آخر لـ `console.log()` للمتغير `learningSentence`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير `learningSentence`.

```js
assert.exists(learningSentence);
```

يجب أن يكون متغير `learningSentence` سلسلة نصية.

```js
assert.isString(learningSentence);
```

يجب أن تعيّن السلسلة النصية `"I love learning!"` إلى المتغير `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

يجب أن تسجل `"Original learning sentence:"` في وحدة التحكم.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

يجب أن تسجل المتغير `learningSentence` في وحدة التحكم.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
