---
id: 68e42de6396f5c49d90d1951
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في درس سابق، تعلمت عن الطريقة `replace()` التي تُستخدم لتحديد سلسلة نصية فرعية داخل سلسلة نصية واستبدالها بقيمة أخرى.

إليك مثالاً:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

تذكر أن السلاسل النصية غير قابلة للتغيير، مما يعني أن النص الأصلي `text` لا يتغير في هذا المثال. بدلاً من ذلك، تُنشأ سلسلة نصية جديدة.

أنشئ متغيرًا جديدًا اسمه `replacedString` وقم بتعيين له القيمة `originalString.replace("cats", "dogs")`.

# --hints--

يجب أن يكون لديك متغير اسمه `replacedString`.

```js
assert.exists(replacedString);
```

يجب أن يكون المتغير `replacedString` سلسلة نصية.

```js
assert.isString(replacedString);
```

يجب أن تعيّن `originalString.replace("cats", "dogs")` إلى المتغير `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
