---
id: 68e448ec4ac4205968c0a805
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

سيكون من الجميل لو تكررت كلمة `"love"` ثلاث مرات في الجملة `"I love learning!"`.

كما تتذكر من درس سابق، يمكنك تكرار سلسلة نصية عددًا محددًا من المرات باستخدام الطريقة `repeat()`. إليك مثالًا:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

بما أن السلاسل النصية غير قابلة للتغيير، فلن تُعدل هذه الطريقة السلسلة الأصلية. ستُرجع سلسلة جديدة بالمحتوى المكرر.

أنشئ متغيرًا اسمه `repeatedLove` وعيّن له القيمة `"love ".repeat(3)`. ثم سجّل المتغير `repeatedLove` في وحدة التحكم لترى النتيجة.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير اسمه `repeatedLove`.

```js
assert.exists(repeatedLove);
```

يجب أن يكون المتغير `repeatedLove` سلسلة نصية.

```js
assert.isString(repeatedLove);
```

يجب أن تعيّن القيمة `"love ".repeat(3)` للمتغير `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

يجب أن تسجل المتغير `repeatedLove` في وحدة التحكم.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
