---
id: 68e4417b7317835208145323
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

بعد ذلك، أنشئ متغيرًا جديدًا اسمه `dogsOnlySentence` وقم بتعيين له القيمة `exampleSentence.replace("cats", "dogs")`.

ثم أسفل ذلك المتغير، سجّل في وحدة التحكم النص `"Replacing all occurrences of cats with dogs:"`. أسفل ذلك الاستدعاء لـ `console.log()`، أضف استدعاءً آخر لـ `console.log()` للمتغير `dogsOnlySentence`.

يجب أن ترى الآن النص `"I love dogs and cats are so much fun!"` مسجلاً في وحدة التحكم. لكن هذا ليس الناتج المطلوب. في الخطوة التالية، ستتعلم كيف تصلحه ليقرأ النص `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن يكون لديك متغير `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

يجب أن يكون المتغير `dogsOnlySentence` من نوع سلسلة نصية.

```js
assert.isString(dogsOnlySentence);
```

يجب أن تعيّن `exampleSentence.replace("cats", "dogs")` إلى المتغير `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

يجب أن تسجل النص `"Replacing all occurrences of cats with dogs:"` في وحدة التحكم.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

يجب أن تسجل المتغير `dogsOnlySentence` في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
