---
id: 68e43e884ef09d4d2ab05631
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

حان الوقت الآن لرؤية النتيجة من استخدام طريقة `replace()`.

ابدأ بإضافة `console.log()` بالقيمة `"After using the replace() method:"`.

ثم أسفل ذلك `console.log()`، أضف `console.log()` آخر مع المتغير `replacedString`.

ألق نظرة على وحدة التحكم ويجب أن ترى أن الجملة الجديدة تقول `"I love dogs."` بدلاً من الأصلية `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل `"After using the replace() method:"` في وحدة التحكم.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

يجب أن تسجل المتغير `replacedString` في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
