---
id: 68e42b62d7093345c6975f22
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

أسفل متغير `originalString`، أضف دالة `console.log()` مع السلسلة النصية `"Original string:"`.

ثم أسفل تلك الدالة `console.log()`، أضف دالة `console.log()` أخرى مع المتغير `originalString`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

يجب أن تسجل السلسلة النصية `"Original string:"` في وحدة التحكم.

```js
assert.equal(getLogs()[0], "Original string:");
```

يجب أن تسجل المتغير `originalString` في وحدة التحكم.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
