---
id: 68e44c305f1f935d949d76bb
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

كما ذُكر في الخطوة السابقة، هناك فراغ إضافي بعد آخر `"love"` في الجملة `"I love love love  learning."`. لإزالة هذا الفراغ الإضافي، يمكنك استخدام الطريقة `trimEnd()`.

في نهاية الطريقة `"love ".repeat(3)`، اربط الطريقة `trimEnd()`. يمكنك ربط الطرق بهذه الطريقة:

```js
.firstMethod().secondMethod()
```

الآن عند فحص وحدة التحكم، يجب أن ترى أن الفراغ الإضافي قد أُزيل.

وبهذا التغيير الأخير، تكون ورشة العمل قد اكتملت!

# --hints--

يجب ربط الطريقة `trimEnd()` بالطريقة `"love ".repeat(3)`.

```js
assert.equal(repeatedLove, "love ".repeat(3).trimEnd());
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--
const repeatedLove = "love ".repeat(3)
--fcc-editable-region--
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```

# --solutions--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3).trimEnd();
console.log(repeatedLove);

const newSentence = `I ${repeatedLove} learning.`;
console.log(newSentence);
```
