---
id: 68e44abea678a65b877864dc
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

سيكون من الجيد رؤية تلك الكلمة المكررة في جملة. أنشئ متغيرًا جديدًا اسمه `newSentence` وقم بتعيين نتيجة ربط الكلمة `"I "` مع المتغير `repeatedLove` ثم ربط السلسلة النصية `" learning."`. بعد ذلك، سجّل المتغير `newSentence` في وحدة التحكم. يمكنك اختيار استخدام عامل `+` أو قوالب حرفية لعملية الربط.

الآن يجب أن ترى النص `"I love love love  learning."`.

*ملاحظة*: هناك مسافة إضافية بعد آخر كلمة `"love"` وهذا صحيح. في الخطوة الأخيرة، ستقوم بقص تلك المسافة الإضافية.

# --hints--

يجب أن يكون لديك متغير اسمه `newSentence`.

```js
assert.exists(newSentence);
```

يجب أن يكون المتغير `newSentence` سلسلة نصية.

```js
assert.isString(newSentence);
```

يجب أن تربط الكلمة `"I "` مع المتغير `repeatedLove` ثم تربط السلسلة النصية `" learning."`. يجب تعيين هذه النتيجة إلى المتغير `newSentence`. انتبه للمسافات.

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

لا يجب أن تعيّن السلسلة النصية الحرفية `"I love love love  learning."` إلى المتغير `newSentence`. يرجى إعادة قراءة التعليمات للطريقة الصحيحة.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
