---
id: 56533eb9ac21ba0edf2244b3
title: Convierte Celsius a Fahrenheit
challengeType: 1
forumTopicId: 16806
dashedName: convert-celsius-to-fahrenheit
---

# --description--

La fórmula para convertir de Celsius a Fahrenheit es la temperatura en Celsius multiplicado por `9/5`, más `32`.

Se dispone una variable `celsius` que representa una temperatura en Celsius. Usa la variable `fahrenheit` ya definida y asígnale la temperatura en Fahrenheit equivalente a la dada en Celsius. Usa la fórmula mencionada arriba para convertir la temperatura de Celsius a Fahrenheit.

# --hints--

`convertCtoF(0)` debería devolver un número

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` debería devolver el valor `-22`

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` debería devolver el valor `14`

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` debería devolver el valor `32`

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` debería devolver el valor `68`

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` debería devolver el valor `86`

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js
function convertCtoF(celsius) {
  let fahrenheit;
  return fahrenheit;
}

convertCtoF(30);
```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
