---
id: 66da326c02141df538f29ba5
title: Construir una página de pago
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un elemento `h1` con el texto `Checkout`.
1. Debes tener dos elementos `section` inmediatamente después del elemento `h1`.
1. Debes tener un elemento `h2` con el texto `Your Cart` dentro de la primera sección.
1. Debes tener una imagen de un artículo en la primera sección con texto alternativo adecuado. Puedes usar esta imagen si lo deseas: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Debes tener un elemento `h2` con el texto `Payment Information` dentro de la segunda sección.
1. Debes tener un elemento `form` dentro de la segunda sección.
1. Debes tener un input con un `id` y `name` de `card-name`, y un `type` de `text` dentro de tu formulario y un `label` asociado a él.
1. Debes tener un input con un `id` y `name` de `card-number`, y un `type` de `text` dentro de tu formulario y una `label` asociada a él.
1. Debes tener al menos dos elementos `input` con el atributo `required`.
1. Debes incluir un elemento `span` con el texto `*` y `aria-hidden` configurado en `true` dentro del elemento `label` para cada entrada requerida, de modo que los campos obligatorios estén indicados visualmente.
1. Debes tener un elemento `p` con un texto de ayuda que explique el formato requerido del número de tarjeta, colocado inmediatamente después del campo de entrada del número de tarjeta. El `p` debe tener un `id` de `card-number-help` y ser referenciado por el campo de entrada del número de tarjeta usando `aria-describedby`.

# --hints--

Debes tener un elemento `h1` con el texto `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Debes tener solo un elemento `h1` en tu página.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debes tener al menos dos elementos `section` después del elemento `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Debes tener un elemento `h2` con el texto `Your Cart` dentro de la primera sección.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Debes tener al menos una imagen con texto alternativo en tu primera sección.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Debes tener un elemento `h2` con el texto `Payment Information` dentro de la segunda sección.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Debes tener un elemento `form` dentro del segundo elemento `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Debes tener un elemento `input` con un `id` y `name` de `card-name` dentro de tu formulario.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Tu campo de nombre de tarjeta debe tener un `type` de `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Debes tener un elemento `input` con un `id` y `name` de `card-number` dentro de tu formulario.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

El campo para ingresar el número de tu tarjeta debe tener un `type` de `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Todos tus elementos `input` que no son de `type` `submit` deben tener un elemento `label` asociado con ellos.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Debes tener al menos dos elementos `input` con el atributo `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Debes incluir un elemento `span` con el texto `*` dentro del elemento `label` para cada entrada requerida.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

Tus elementos `span` deben tener `aria-hidden` configurado en `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Debes tener un elemento `p` con un `id` de `card-number-help` inmediatamente después del campo de número de tarjeta.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

El texto de ayuda de tu número de tarjeta no debe estar vacío.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

El campo de entrada de número de tarjeta debe tener `aria-describedby` configurado en `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
