---
id: 587d7fa5367417b2b2512bbd
title: 다른 요소에 CSS 스타일 한 세트 확장하기
challengeType: 0
forumTopicId: 301456
dashedName: extend-one-set-of-css-styles-to-another-element
---

# --description--

Sass에는 한 요소의 CSS 규칙을 쉽게 빌려와 다른 요소에서 확장할 수 있는 `extend` 기능이 있습니다.

예를 들어, 아래 CSS 규칙 블록은 `.panel` 클래스를 스타일링합니다. `background-color`, `height`, `border`가 포함되어 있습니다.

```scss
.panel{
  background-color: red;
  height: 70px;
  border: 2px solid green;
}
```

이제 `.big-panel`라는 또 다른 패널이 필요합니다. `.panel`과 동일한 기본 속성을 가지면서 `width`과 `font-size`도 필요합니다. `.panel`에서 초기 CSS 규칙을 복사해 붙여넣을 수도 있지만, 패널 유형이 많아질수록 코드가 반복적이 됩니다. `extend` 지시문은 한 요소에 작성된 규칙을 재사용하고 다른 요소에 추가 스타일을 더하는 간단한 방법입니다:

```scss
.big-panel{
  @extend .panel;
  width: 150px;
  font-size: 2em;
}
```

`.big-panel`은 `.panel`와 동일한 속성에 새로운 스타일이 추가됩니다.

# --instructions--

`.info-important` 클래스를 만들어 `.info`를 확장하고 `background-color`를 magenta로 설정하세요.

# --hints--

`info-important` 클래스는 `background-color`가 `magenta`로 설정되어야 합니다.

```js
assert(
  code.match(
    /\.info-important\s*?{[\s\S]*background-color\s*?:\s*?magenta\s*?;[\s\S]*}/gi
  )
);
```

`info-important` 클래스는 `@extend`를 사용해 `info` 클래스의 스타일을 상속해야 합니다.

```js
assert(
  code.match(/\.info-important\s*?{[\s\S]*@extend\s*?.info\s*?;[\s\S]*/gi)
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }




</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  h3{
    text-align: center;
  }
  .info{
    width: 200px;
    border: 1px solid black;
    margin: 0 auto;
  }
  .info-important{
    @extend .info;
    background-color: magenta;
  }



</style>
<h3>Posts</h3>
<div class="info-important">
  <p>This is an important post. It should extend the class ".info" and have its own CSS styles.</p>
</div>

<div class="info">
  <p>This is a simple post. It has basic styling and can be extended for other uses.</p>
</div>
```
