---
id: 587d7dbd367417b2b2512bb6
title: 믹스인으로 재사용 가능한 CSS 만들기
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

Sass에서 <dfn>믹스인</dfn>은 스타일 시트 전체에서 재사용할 수 있는 CSS 선언 그룹입니다. 정의는 `@mixin` @규칙으로 시작하며, 그 뒤에 사용자 지정 이름이 옵니다. 믹스인은 `@include` @규칙을 사용해 적용합니다.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

컴파일 결과:

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

믹스인은 인수를 받을 수도 있어서 동작을 맞춤 설정할 수 있습니다. 믹스인을 사용할 때 인수는 필수입니다.

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

매개변수에 기본값을 지정하면 인수를 선택 사항으로 만들 수 있습니다.

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

`shape`라는 이름의 믹스인을 작성하고 `$w`, `$h`, `$bg-color` 세 개의 매개변수를 만드세요.

`shape` 믹스인을 사용해 `#square` 요소에 너비와 높이를 `50px`로, 배경색을 `red`로 지정하세요. `#rect-a` 요소에는 너비 `100px`, 높이 `50px`, 배경색 `blue`을 추가하세요. 마지막으로 `#rect-b` 요소에는 너비 `50px`, 높이 `100px`, 배경색 `orange`을 추가하세요.

# --hints--

`shape`라는 이름의 믹스인을 세 개의 매개변수 `$w`, `$h`, `$bg-color`와 함께 선언하세요.

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

믹스인에 `width` 속성을 포함하고 `$w` 매개변수를 사용하세요.

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

믹스인에 `height` 속성을 포함하고 `$h` 매개변수를 사용하세요.

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

믹스인에 `background-color` 속성을 포함하고 `$bg-color` 매개변수를 사용하세요.

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

`#square` 선택자 안의 스타일을 `shape` 믹스인을 `@include` 키워드로 호출하는 코드로 바꾸세요. 너비와 높이를 `50px`로, 배경색을 `red`로 설정하세요.

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

`#rect-a` 선택자 안의 스타일을 `shape` 믹스인을 `@include` 키워드로 호출하는 코드로 바꾸세요. 너비 `100px`, 높이 `50px`, 배경색 `blue`을 설정하세요.

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

`#rect-b` 선택자 안의 스타일을 `shape` 믹스인을 `@include` 키워드로 호출하는 코드로 바꾸세요. 너비 `50px`, 높이 `100px`, 배경색 `orange`을 설정하세요.

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
