---
id: 587d7dbf367417b2b2512bbc
title: 파셜로 스타일을 더 작은 조각으로 나누기
challengeType: 0
forumTopicId: 301459
dashedName: split-your-styles-into-smaller-chunks-with-partials
---

# --description--

<dfn>파셜</dfn>은 Sass에서 CSS 코드 조각을 담는 별도의 파일입니다. 이 파일들은 다른 Sass 파일에서 가져와 사용합니다. 비슷한 코드를 모듈로 묶어 정리하는 좋은 방법입니다.

파셜 이름은 밑줄(`_`) 문자로 시작하는데, 이는 Sass에 이것이 작은 CSS 조각이며 CSS 파일로 변환하지 말라는 뜻입니다. 또한 Sass 파일은 `.scss` 확장자로 끝납니다. 파셜의 코드를 다른 Sass 파일에 가져오려면 `@import` 지시어를 사용하세요.

예를 들어, 모든 믹스인이 "\_mixins.scss"라는 파셜에 저장되어 있고, "main.scss" 파일에서 필요하다면, 메인 파일에서 이렇게 사용합니다:

```scss
@import 'mixins'
```

밑줄과 파일 확장자는 `import` 문장에 포함하지 않아도 됩니다 — Sass가 파셜임을 인식합니다. 파셜이 파일에 가져와지면 모든 변수, 믹스인, 기타 코드가 사용 가능해집니다.

# --instructions--

main.scss 파일에 `@import`라는 파셜을 가져오는 `_variables.scss` 문장을 작성하세요.

# --hints--

코드는 `@import` 지시어를 사용해야 하며, 파일 이름에 밑줄은 포함하지 마세요.

```js
assert(code.match(/@import\s+?('|")variables\1/gi));
```

# --seed--

## --seed-contents--

```html
<!-- The main.scss file -->
```

# --solutions--

```html
@import 'variables'
```
