---
id: 587d7dbe367417b2b2512bb9
title: "@for를 사용해 Sass 반복문 만들기"
challengeType: 0
forumTopicId: 301462
dashedName: use-for-to-create-a-sass-loop
---

# --description--

`@for` 지시문은 JavaScript의 `for` 반복문과 매우 비슷하게 스타일을 반복해서 추가합니다.

`@for`는 두 가지 방식으로 사용됩니다: "start through end" 또는 "start to end". 주요 차이점은 "start **to** end"가 끝 숫자를 개수에 포함하지 않고, "start **through** end"는 끝 숫자를 개수에 포함한다는 점입니다.

다음은 start **through** end 예시입니다:

```scss
@for $i from 1 through 12 {
  .col-#{$i} { width: 100%/12 * $i; }
}
```

`#{$i}` 부분은 변수(`i`)와 텍스트를 결합해 문자열을 만드는 구문입니다. Sass 파일이 CSS로 변환되면 다음과 같이 보입니다:

```scss
.col-1 {
  width: 8.33333%;
}

.col-2 {
  width: 16.66667%;
}

...

.col-12 {
  width: 100%;
}
```

이 방법은 그리드 레이아웃을 만들 때 매우 강력합니다. 이제 CSS 클래스 형태로 열 너비를 12가지 옵션으로 사용할 수 있습니다.

# --instructions--

변수 `@for`가 1 **to** 6까지 가는 `$j` 지시문을 작성하세요.

각 인덱스에 15px을 곱한 값을 `.text-1`로 설정하는 `.text-5`부터 `font-size`까지 5개의 클래스를 만들어야 합니다.

# --hints--

코드는 `@for` 지시문을 사용해야 합니다.

```js
assert(code.match(/@for /g));
```

`.text-1` 클래스는 `font-size`가 15px이어야 합니다.

```js
assert($('.text-1').css('font-size') == '15px');
```

`.text-2` 클래스는 `font-size`가 30px이어야 합니다.

```js
assert($('.text-2').css('font-size') == '30px');
```

`.text-3` 클래스는 `font-size`가 45px이어야 합니다.

```js
assert($('.text-3').css('font-size') == '45px');
```

`.text-4` 클래스는 `font-size`가 60px이어야 합니다.

```js
assert($('.text-4').css('font-size') == '60px');
```

`.text-5` 클래스는 `font-size`가 75px이어야 합니다.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>

@for $i from 1 through 5 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

---

```html
<style type='text/scss'>

@for $i from 1 to 6 {
  .text-#{$i} { font-size: 15px * $i; }
}

</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
