---
id: 587d7dbd367417b2b2512bb4
title: Sass 변수로 데이터 저장하기
challengeType: 0
forumTopicId: 301460
dashedName: store-data-with-sass-variables
---

# --description--

Sass가 CSS와 다른 기능 중 하나는 변수를 사용한다는 점입니다. 변수는 선언되어 데이터를 저장하는 데 사용되며, JavaScript와 비슷합니다.

JavaScript에서는 변수를 `let`와 `const` 키워드를 사용해 정의합니다. Sass에서는 변수 이름 앞에 `$`가 붙습니다.

몇 가지 예는 다음과 같습니다:

```scss
$main-fonts: Arial, sans-serif;
$headings-color: green;
```

그리고 변수를 사용하는 방법은 다음과 같습니다:

```scss
h1 {
  font-family: $main-fonts;
  color: $headings-color;
}
```

변수가 유용한 예 중 하나는 여러 요소가 같은 색상을 가져야 할 때입니다. 그 색상을 변경할 경우, 코드를 수정해야 하는 곳은 변수 값뿐입니다.

# --instructions--

변수 `$text-color`를 선언하고 `red`로 설정하세요. 그런 다음 `color`와 `.blog-post`의 `h2` 속성 값을 `$text-color` 변수로 변경하세요.

# --hints--

`$text-color`에 대해 값이 `red`인 Sass 변수가 선언되어 있어야 합니다.

```js
assert(code.match(/\$text-color\s*:\s*?red\s*;/g));
```

`$text-color`와 `color` 항목의 `.blog-post`를 변경하기 위해 `h2` 변수를 사용해야 합니다.

```js
assert(code.match(/color\s*:\s*\$text-color\s*;?/g));
```

`.blog-post` 요소는 `color`가 빨간색이어야 합니다.

```js
assert($('.blog-post').css('color') == 'rgb(255, 0, 0)');
```

`h2` 요소들은 `color`가 빨간색이어야 합니다.

```js
assert($('h2').css('color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>


  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: red;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  $text-color: red;

  .header{
    text-align: center;
  }
  .blog-post, h2 {
    color: $text-color;
  }
</style>

<h1 class="header">Learn Sass</h1>
<div class="blog-post">
  <h2>Some random title</h2>
  <p>This is a paragraph with some random text in it</p>
</div>
<div class="blog-post">
  <h2>Header #2</h2>
  <p>Here is some more random text.</p>
</div>
<div class="blog-post">
  <h2>Here is another header</h2>
  <p>Even more random text within a paragraph</p>
</div>
```
