---
id: 587d7dbd367417b2b2512bb5
title: Sass로 CSS 중첩하기
challengeType: 0
forumTopicId: 301457
dashedName: nest-css-with-sass
---

# --description--

Sass는 CSS 규칙을 중첩할 수 있게 해주며, 이는 스타일 시트를 체계적으로 정리하는 데 유용한 방법입니다.

보통 각 요소는 스타일을 지정하기 위해 다른 줄에 타겟팅됩니다. 예를 들어 다음과 같습니다:

```scss
article {
  height: 200px;
}

article p {
  color: white;
}

article ul {
  color: blue;
}
```

대규모 프로젝트에서는 CSS 파일에 많은 줄과 규칙이 포함됩니다. 이때 중첩을 사용하면 자식 스타일 규칙을 해당 부모 요소 안에 배치하여 코드를 정리하는 데 도움이 됩니다:

```scss
article {
  height: 200px;

  p {
    color: white;
  }

  ul {
    color: blue;
  }
}

```

# --instructions--

위에 나온 중첩 기법을 사용해 `.blog-post` 요소의 두 자식에 대한 CSS 규칙을 다시 정리하세요. 테스트를 위해 `h1` 요소가 `p` 요소보다 먼저 와야 합니다.

# --hints--

코드를 다시 정리하여 `h1`과 `p`가 `.blog-post` 부모 요소 안에 중첩되도록 하세요.

```js
assert(
  code.match(
    /\.blog-post\s*?{\s*?h1\s*?{\s*?text-align:\s*?center;\s*?color:\s*?blue;\s*?}\s*?p\s*?{\s*?font-size:\s*?20px;\s*?}\s*?}/gi
  )
);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
  .blog-post {

  }
  h1 {
    text-align: center;
    color: blue;
  }
  p {
    font-size: 20px;
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```

# --solutions--

```html
<style type='text/scss'>
  .blog-post {
    h1 {
      text-align: center;
      color: blue;
    }
    p {
      font-size: 20px;
    }
  }
</style>

<div class="blog-post">
  <h1>Blog Title</h1>
  <p>This is a paragraph</p>
</div>
```
