---
id: 587d7dbf367417b2b2512bbb
title: 조건이 충족될 때까지 @while로 스타일 적용하기
challengeType: 0
forumTopicId: 301454
dashedName: apply-a-style-until-a-condition-is-met-with-while
---

# --description--

`@while` 지시문은 JavaScript `while` 반복문과 유사한 기능을 가진 옵션입니다. 조건이 충족될 때까지 CSS 규칙을 생성합니다.

`@for` 도전 과제에서는 간단한 그리드 시스템을 만드는 예를 제시했습니다. 이것은 `@while`에도 적용할 수 있습니다.

```scss
$x: 1;
@while $x < 13 {
  .col-#{$x} { width: 100%/12 * $x;}
  $x: $x + 1;
}
```

먼저, 변수 `$x`를 정의하고 1로 설정하세요. 다음으로, `@while` 지시문을 사용해 `$x`가 13 미만인 *동안* 그리드 시스템을 만드세요. `width`에 대한 CSS 규칙을 설정한 후, 무한 반복을 피하기 위해 `$x`를 1씩 증가시키세요.

# --instructions--

`@while`를 사용해 서로 다른 `font-sizes`을 가진 일련의 클래스를 만드세요.

`text-1`부터 `text-5`까지 5개의 서로 다른 클래스가 있어야 합니다. 그런 다음 `font-size`을 현재 인덱스 번호에 `15px`를 곱한 값으로 설정하세요. 무한 반복이 발생하지 않도록 주의하세요!

# --hints--

코드에 `@while` 지시문을 사용해야 합니다.

```js
assert(code.match(/@while /g));
```

코드에 인덱스 변수를 사용하며, 인덱스는 1부터 시작해야 합니다.

```js
assert(code.match(/\$.*:\s*?1;/gi));
```

카운터 변수를 증가시켜야 합니다.

```js
assert(code.match(/\$(.*)\s*?:\s*\$\1\s*\+\s*1\s*;/gi));
```

`.text-1` 클래스는 `font-size`이 `15px`이어야 합니다.

```js
assert($('.text-1').css('font-size') == '15px');
```

`.text-2` 클래스는 `font-size`이 `30px`이어야 합니다.

```js
assert($('.text-2').css('font-size') == '30px');
```

`.text-3` 클래스는 `font-size`이 `45px`이어야 합니다.

```js
assert($('.text-3').css('font-size') == '45px');
```

`.text-4` 클래스는 `font-size`이 `60px`이어야 합니다.

```js
assert($('.text-4').css('font-size') == '60px');
```

`.text-5` 클래스는 `font-size`이 `75px`이어야 합니다.

```js
assert($('.text-5').css('font-size') == '75px');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```

# --solutions--

```html
<style type='text/scss'>
  $x: 1;
  @while $x < 6 {
    .text-#{$x}{
      font-size: 15px * $x;
    }
    $x: $x + 1;
  }
</style>

<p class="text-1">Hello</p>
<p class="text-2">Hello</p>
<p class="text-3">Hello</p>
<p class="text-4">Hello</p>
<p class="text-5">Hello</p>
```
