---
id: 587d7dbf367417b2b2512bba
title: 리스트 항목에 @each를 사용해 반복하기
challengeType: 0
forumTopicId: 301461
dashedName: use-each-to-map-over-items-in-a-list
---

# --description--

마지막 도전 과제에서는 `@for` 지시어가 시작값과 끝값을 사용해 정해진 횟수만큼 반복하는 방법을 보여주었습니다. Sass는 리스트나 맵의 각 항목을 반복하는 `@each` 지시어도 제공합니다. 각 반복마다 변수에 리스트나 맵의 현재 값이 할당됩니다.

```scss
@each $color in blue, red, green {
  .#{$color}-text {color: $color;}
}
```

맵은 약간 다른 구문을 가집니다. 예시는 다음과 같습니다:

```scss
$colors: (color1: blue, color2: red, color3: green);

@each $key, $color in $colors {
  .#{$color}-text {color: $color;}
}
```

맵의 키를 참조하려면 `$key` 변수가 필요하다는 점에 주의하세요. 그렇지 않으면 컴파일된 CSS에 `color1`, `color2`... 같은 값이 들어가게 됩니다. 위 두 코드 예시는 다음 CSS로 변환됩니다:

```scss
.blue-text {
  color: blue;
}

.red-text {
  color: red;
}

.green-text {
  color: green;
}
```

# --instructions--

리스트 `@each`를 순회하는 `blue, black, red` 지시어를 작성하고, 각 변수를 `.color-bg` 클래스에 할당하세요. 여기서 `color` 부분은 각 항목에 맞는 색상으로 바뀝니다. 각 클래스는 `background-color`도 해당 색상으로 설정해야 합니다.

# --hints--

코드는 `@each` 지시어를 사용해야 합니다.

```js
assert(code.match(/@each /g));
```

`.blue-bg` 클래스는 `background-color`가 파란색이어야 합니다.

```js
assert($('.blue-bg').css('background-color') == 'rgb(0, 0, 255)');
```

`.black-bg` 클래스는 `background-color`가 검은색이어야 합니다.

```js
assert($('.black-bg').css('background-color') == 'rgb(0, 0, 0)');
```

`.red-bg` 클래스는 `background-color`가 빨간색이어야 합니다.

```js
assert($('.red-bg').css('background-color') == 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

# --solutions--

```html
<style type='text/scss'>

  @each $color in blue, black, red {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```

---

```html
<style type='text/scss'>

  $colors: (color1: blue, color2: black, color3: red);

  @each $key, $color in $colors {
    .#{$color}-bg {background-color: $color;}
  }

  div {
    height: 200px;
    width: 200px;
  }
</style>

<div class="blue-bg"></div>
<div class="black-bg"></div>
<div class="red-bg"></div>
```
