---
id: 587d7dbe367417b2b2512bb8
title: 스타일에 논리를 추가하기 위해 @if와 @else 사용하기
challengeType: 0
forumTopicId: 301463
dashedName: use-if-and-else-to-add-logic-to-your-styles
---

# --description--

Sass에서 `@if` 지시문은 특정 경우를 테스트할 때 유용합니다 - JavaScript의 `if` 문과 똑같이 작동합니다.

```scss
@mixin make-bold($bool) {
  @if $bool == true {
    font-weight: bold;
  }
}
```

그리고 JavaScript에서처럼 `@else if`와 `@else` 지시문은 더 많은 조건을 테스트합니다:

```scss
@mixin text-effect($val) {
  @if $val == danger {
    color: red;
  }
  @else if $val == alert {
    color: yellow;
  }
  @else if $val == success {
    color: green;
  }
  @else {
    color: black;
  }
}
```

# --instructions--

`border-stroke`라는 믹스인을 선언하고 `$val`라는 매개변수를 받으세요. 믹스인은 `@if`, `@else if`, `@else` 지시문을 사용해 다음 조건들을 검사해야 합니다:

```scss
light - 1px solid black
medium - 3px solid black
heavy - 6px solid black
```

만약 `$val` 매개변수 값이 `light`, `medium`, `heavy`가 아니라면, `border` 속성은 `none`로 설정되어야 합니다.

# --hints--

`border-stroke`라는 이름의 믹스인을 선언하고 `$val`라는 매개변수를 포함해야 합니다.

```js
assert(code.match(/@mixin\s+?border-stroke\s*?\(\s*?\$val\s*?\)\s*?{/gi));
```

믹스인 안에 `@if` 문이 있어 `$val`이 `light`인지 확인하고 `border`을 `1px solid black`로 설정해야 합니다.

```js
assert(
  code.match(
    /@if\s+?\$val\s*?===?\s*?light\s*?{\s*?border\s*?:\s*?1px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

믹스인 안에 `@else if` 문이 있어 `$val`이 `medium`인지 확인하고 `border`를 `3px solid black`로 설정해야 합니다.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?medium\s*?{\s*?border\s*?:\s*?3px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

믹스인 안에 `@else if` 문이 있어 `$val`이 `heavy`인지 확인하고 `border`을 `6px solid black`로 설정해야 합니다.

```js
assert(
  code.match(
    /@else\s+?if\s+?\$val\s*?===?\s*?heavy\s*?{\s*?border\s*?:\s*?6px\s+?solid\s+?black\s*?;\s*?}/gi
  )
);
```

믹스인 안에 `@else` 문이 있어 `border`을 `none`로 설정해야 합니다.

```js
assert(code.match(/@else\s*?{\s*?border\s*?:\s*?none\s*?;\s*?}/gi));
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>



  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```

# --solutions--

```html
<style type='text/scss'>
  @mixin border-stroke($val) {
    @if $val == light {
      border: 1px solid black;
    }
    @else if $val == medium {
      border: 3px solid black;
    }
    @else if $val == heavy {
      border: 6px solid black;
    }
    @else {
      border: none;
    }
  }


  #box {
    width: 150px;
    height: 150px;
    background-color: red;
    @include border-stroke(medium);
  }
</style>

<div id="box"></div>
```
