---
id: 66b102ee0bdbad34a9f85ba0
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você vai continuar aprendendo sobre strings criando um professor chatbot.

Para começar, adicione uma instrução de `console`, com a mensagem `"Hi there!"`.

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve ter uma instrução de `console`.

```js
assert.isAbove(logSpy.calls.length, 0);
```

A instrução de console deve exibir a mensagem `"Hi there!"`.

```js
assert.deepInclude(logSpy.calls, ["Hi there!"]);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
