---
id: 66b102ee0bdbad34a9f85ba0
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você vai continuar aprendendo sobre strings criando um professor chatbot. 

Para começar, adicione uma instrução de `console`, com a mensagem `"Hi there!"`.

# --hints--

Você deve ter uma instrução de `console`.

```js
assert.match(code, /console\.log(.*)/);
```

A instrução de console deve exibir a mensagem `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
