---
id: 66b5928b86e907fdfbf34e56
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora é hora de definir o nome do bot.

Crie uma variável chamada `botName` e atribua o valor de string `"teacherBot"`.

# --hints--

Você deve ter uma variável chamada `botName`.

```js
assert.isNotNull(botName);
```

A variável `botName` deve ser uma string.

```js
assert.isString(botName);
```

A variável `botName` deve ter o valor `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
