---
id: 66b59829dba144ff1351220f
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora é hora de criar uma saudação usando a variável `botName`.

Nas lições anteriores, você aprendeu como concatenar strings usando template literals assim:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Comece criando uma variável chamada `greeting`.

Em seguida, usando a sintaxe de template literals, atribua uma string que diga `My name is`, seguida pela variável `botName` e terminando com um ponto (`.`).

Finalmente, registre a variável `greeting` no console.

# --hints--

Você deve ter uma variável chamada `greeting`.

```js
assert.isNotNull(greeting);
```

Você deve atribuir uma string à variável `greeting`.

```js
assert.isString(greeting);
```

Você deve usar template literals para concatenar a string `My name is` com a variável `botName` seguida por um ponto (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Você deve ter uma instrução de `console` que registre a variável `greeting` no console.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
