---
id: 66b6fdb76441c738719039fa
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Nas lições anteriores, você aprendeu como acessar o último caractere em uma string assim:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` sempre mostrará o último número de índice de uma string.

Crie uma variável chamada `lastCharacter` e atribua a ela o valor do último caractere na variável `subject`.

Em seguida, registre o valor da variável `lastCharacter` no console.

# --hints--

Você deve ter uma variável chamada `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Você deve atribuir o valor do último caractere na variável `subject` à variável `lastCharacter`. Consulte o exemplo se precisar de ajuda.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Você deve registrar o valor da variável `lastCharacter` no `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
