---
id: 66b6e39031393e30f2c48d0f
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Para obter o comprimento de uma string, você pode usar a propriedade <dfn>length</dfn>.Essa propriedade retorna o número de caracteres em uma string.

Aqui está um exemplo:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Comece usando a propriedade `length` para obter o comprimento da string `subject` e registre esse valor no console.

# --hints--


Você deve usar a propriedade `length` para obter o comprimento da string `subject`.

```js
assert.match(code, /subject\.length/);
```

Você deve registrar o comprimento da string `subject` no console.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
