---
id: 66b6ffb42b88e33943788abf
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Para a última parte do workshop, você vai revisar como encontrar a posição do índice de uma substring em uma string. Lembre-se de que uma substring é uma parte de uma string.

Comece criando uma variável chamada `learningIsFunSentence` e atribua a ela o valor da string `"Learning is fun."`.

# --hints--

Você deve ter uma variável chamada `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

`learningIsFunSentence` deve conter o valor de uma string.

```js
assert.isString(learningIsFunSentence);
```

Você deve atribuir o valor `"Learning is fun."` à variável `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
