---
id: 66b6e62423e8a031d6c1f03d
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Agora é hora de obter o comprimento da string `topic`.

Você pode usar template literals dentro de instruções de `console` assim:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Comece exibindo a mensagem `Here is an example of using the length property on the word [topic].` para o console. 

Lembre-se de substituir `[topic]` pela variável `topic` e de usar a sintaxe apropriada de template literals, como você fez nos passos anteriores.

Em seguida, adicione uma segunda instrução de `console.log` que exiba o comprimento da string `topic` no console.

# --hints--

Você deve usar a sintaxe de template literals para exibir a mensagem `Here is an example of using the length property on the word [topic].` no `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Você deve ter uma segunda instrução de `console` para exibir o comprimento da variável `topic` no `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
