---
id: 66b6f586767a1534f3097353
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora é hora de acessar a segunda letra da variável `subject`.

Comece adicionando uma instrução de `console` que produz a mensagem `Here is an example of accessing the second letter in the word [subject].` Lembre-se de substituir `[subject]` pelo valor real da variável `subject` e use a sintaxe correta de template literals.

Em seguida, adicione outra instrução de `console` que exiba a segunda letra da variável `subject` usando notação de colchetes e o número de índice correto.

# --hints--

A instrução de `console` deve produzir a mensagem `Here is an example of accessing the second letter in the word [subject].`. 

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Você deve ter outra instrução de `console` que exiba a segunda letra da variável `subject` usando o índice `1` com notação de colchetes.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
