---
id: 688141e05e367de0e3ef7635
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

此刻，`currentContent` 被初始化为一个空字串的值。但期望的结果是让 `currentContent` 保存当前笔记文本的值。

在之前的课程中，你学习了如何像这样处理 `DOMContentLoaded` 事件：

```js
window.addEventListener("DOMContentLoaded", () => {
  // Do stuff here
});
```

当超文本标记语言文档中的所有内容都已装载并解析后，你将能够安全地访问 `noteEl`。

首先为 `window` 对象附加一个 `addEventListener` 方法。`addEventListener` 方法的第一个参数应为 `"DOMContentLoaded"` 事件。第二个参数应为一个箭头函数。在箭头函数的体内，将 `noteEl.textContent` 的值赋值给 `currentContent`。

# --hints--

你应该为 `window` 对象添加一个 `addEventListener` 方法。

```js
assert.match(code, /window\.addEventListener/);
```

你的事件监听器应监听 `"DOMContentLoaded"` 事件。

```js
assert.match(code, /window\.addEventListener\(['"]DOMContentLoaded['"]\,/);
```

在箭头函数的体内，将 `noteEl.textContent` 的值赋值给 `currentContent`。

```js
const event = new Event("DOMContentLoaded");
window.dispatchEvent(event);

assert.equal(currentContent, noteEl.textContent);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English.

      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.

      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>

    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--

--fcc-editable-region--
```
