---
id: 6881417fdb4032df98134a89
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

接下来的工作坊步骤将涉及为动态显示已储存消息和将当前笔记日志到控制台添加代码。

首先创建一个名为 `noteEl` 的变量，并将查询文档中 `id` 为 `note` 的元素的结果赋值给它。

然后创建另一个名为 `statusEl` 的变量，并将查询文档中 `id` 为 `status` 的元素的结果赋值给它。

# --hints--

你应该有一个 `noteEl` 变量，并将其赋值为查询文档中 `id` 为 `note` 的元素的结果。

```js
assert.equal(noteEl, document.getElementById("note"));
``` 

你应该有一个 `statusEl` 变量，并将其赋值为查询文档中 `id` 为 `status` 的元素的结果。

```js
assert.equal(statusEl, document.getElementById("status"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
--fcc-editable-region--

--fcc-editable-region--
```
