---
id: 68814221fc4752e17b9696bd
title: 步骤 8
challengeType: 0
dashedName: step-8
---

# --description--

当用户点击笔记、进行编辑然后点击笔记外部时，你需要一种方法来获取新编辑的版本。你可以通过使用 `blur` 事件来实现。当元素失去焦点时，该事件会触发。以下是一个示例：

```js
element.addEventListener("blur", () => {
  // do something here
});
```

为你的 `noteEl` 变量添加一个 `addEventListener` 方法。第一个参数应为 `"blur"` 事件，第二个参数应为一个箭头函数。在该箭头函数的体内，使用 `const` 创建一个名为 `newContent` 的新变量，并赋值为 `noteEl.innerHTML`。 

这里使用 `innerHTML` 来帮助保留注释的间距和格式。 

# --hints--

你应该为 `noteEl` 变量添加一个 `addEventListener` 方法。

```js
assert.match(code, /noteEl\.addEventListener/);
```

你的 `addEventListener` 方法的第一个参数应为 `"blur"` 事件。

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],/);
```

在箭头函数的体内，你应该使用 `const` 创建一个名为 `newContent` 的新变量，并赋值为 `noteEl.innerHTML`。

```js
assert.match(code, /noteEl\.addEventListener\(['"]blur['"],\s*\(\)\s*=>\s*{[^}]*const\s+newContent\s*=\s*noteEl\.innerHTML[^}]*}\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Note taking app</title>
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <p class="helper-text">Click or tap on the card to edit your note.</p>

    <div id="note" class="note" contenteditable="true" aria-label="Note editor">
      Many languages have words that carry meanings so specific or culturally rooted that they can't be neatly translated into English. 
        
      One example is the Japanese word "tsundoku", which refers to the habit of acquiring books and letting them pile up unread, something many book lovers can relate to. Another is the Portuguese word "saudade", describing a deep, bittersweet longing for something or someone that is absent. Meanwhile, the French word "Dépaysement" captures the disorienting yet exciting feeling of being in a new place, far from home.
        
      These unique words remind us that language is more than vocabulary: it's a window into the values, habits, and emotions of the cultures that create it.
    </div>

    <div id="status" aria-live="polite"></div>
    
    <script src="script.js"></script>
  </body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 2em;
  max-width: 700px;
  background-color: #f5f5f5;
}

.note {
  background-color: #ffffff;
  border: 1px solid #ddd;
  border-radius: 4px;
  box-shadow: 0 2px 5px rgba(0, 0, 0, 0.1);
  padding: 1.5em;
  margin-bottom: 1em;
  line-height: 1.5;
  min-height: 250px;
  font-size: 16px;
  /* This is needed to preserve line breaks in the div */
  white-space: pre-wrap;
}

.note[contenteditable="true"] {
  caret-color: black;
}

.note:hover {
  background-color: #fff;
  box-shadow: 0 0 5px rgba(0,0,0,0.2);
}

.helper-text {
  font-size: 0.9rem;
  color: #666;
  margin-top: 0.5em;
  user-select: none;
  font-style: italic;
}

#status {
  color: #00471b;
  padding: 0 1em;
}
```

```js
const noteEl = document.getElementById("note");
const statusEl = document.getElementById("status");

let currentContent = "";

--fcc-editable-region--
  
--fcc-editable-region--

window.addEventListener("DOMContentLoaded", () => {
  currentContent = noteEl.textContent;
});
```
