---
id: 6853e58d52ec64349a5760ca
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви реалізуєте алгоритм пошуку найкоротшого шляху. Ви напишете Python-функцію, яка обчислює найкоротший шлях між вузлами в графі, а також повертає пройдений маршрут.

Наприклад, якщо є граф, де міста з’єднані дорогами з різною довжиною, алгоритм знайде найкоротший маршрут від одного міста до іншого. Якщо ви хочете поїхати з міста A до міста D, алгоритм може визначити, що шлях A ⇨ B ⇨ C ⇨ D (загалом 15 км) коротший, ніж прямий шлях A ⇨ D (20 км).

Щоб почати, визначте змінну з ім’ям `INF` і призначте їй значення `float('inf')`, яке представляє позитивну нескінченність. Пізніше ви використаєте це, щоб позначити нескінченну відстань між двома вузлами.

# --hints--

Визначте змінну `INF` на початку вашого файлу.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Призначте змінній `float('inf')` значення `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
