---
id: 686253f1b051998ad4904e3e
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Тепер ви створите головну функцію, яка прийматиме три параметри: матрицю суміжності, початковий вузол і необов’язковий цільовий вузол.

Створіть функцію з назвою `shortest_path`, яка прийматиме три параметри: `matrix`, `start_node` і `target_node`. Призначте `None` як значення за замовчуванням для `target_node`.

Параметр `target_node` за замовчуванням має значення `None`, що означає, що якщо цільовий вузол не вказано, функція має обчислити найкоротші шляхи від початкового вузла до всіх інших вузлів графа.

Поки що додайте інструкцію `pass` всередині тіла функції.

# --hints--

Ви маєте створити функцію з назвою `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

Функція має приймати три параметри: `matrix`, `start_node` і `target_node`, причому `target_node` має значення за замовчуванням `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
