---
id: 686251fb49ebc489ba60ccb5
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Вам потрібно створити двовимірний список, щоб представити матрицю суміжності графа. Ця матриця буде використана для представлення ваг ребер між вузлами графа.

Створіть змінну з назвою `adj_matrix` і призначте їй двовимірний список, який представляє граф із такими вагами:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Ви повинні створити змінну з назвою `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` має бути двовимірним списком, що містить надані значення.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
