---
id: 686b82af323098e7a5203bf4
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Під час виконання алгоритму потрібно відстежувати, які вузли ви вже відвідали, щоб не обробляти їх більше одного разу.

Для цього створіть список з назвою `visited` і ініціалізуйте його `False` для кожного вузла.

# --hints--

У вас має бути список з назвою `visited` всередині функції `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("visited")`
        )
      );
    }
});
```

Потрібно ініціалізувати список `visited` значенням `False` для кожного вузла, використовуючи `[False] * n`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("visited").is_equivalent("visited = [False] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]

--fcc-editable-region--
```
