---
id: 686b7848fbaa41e0c2a7efef
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер, коли ваш список `distances` ініціалізовано, потрібно оновити відстань для початкового вузла.

Оскільки відстань від початкового вузла до самого себе завжди `0`, встановіть значення за індексом `start_node` у списку `distances` рівним `0`.

# --hints--

Ви повинні встановити `distances[start_node]` у `0`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_stmt("distances[start_node] = 0")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n

--fcc-editable-region--
```
