---
id: 686b78bc13d518e1267448e4
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Окрім відстеження відстаней, вам також потрібно зберігати інформацію про фактичні шляхи, якими досягають кожного вузла.

Ви створите список, де кожен елемент зберігатиме шлях, пройдений до відповідного вузла. Спочатку шлях кожного вузла міститиме лише його самого.

Спискові включення дають лаконічний спосіб створення списків. Наприклад:

```py
[x * 2 for x in range(3)]
```

Створіть змінну з ім’ям `paths` і ініціалізуйте її за допомогою спискового включення, яке створює список, що містить `[node_no]` для кожного `node_no` у `range(n)`.

# --hints--

У вас має бути змінна з ім’ям `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Ви маєте ініціалізувати `paths` за допомогою спискового включення.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Ваше спискове включення має ітеруватися над `node_no` за допомогою `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Ваше спискове включення має обчислювати `[node_no]` для кожного `node_no` у `range(n)`.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0

--fcc-editable-region--
```
