---
id: 68760671d950767b055bf2f9
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Тепер потрібно перевірити кожен вузол, щоб знайти той, у якого найменша відома відстань і який ще не був відвіданий.

Для цього додайте `for` цикл всередину головного циклу.

Цикл має ітерувати по кожному `node_no` у `range(n)`, де `n` — це число вузлів у графі. Наразі додайте `pass` як тіло циклу.

# --hints--

У вас має бути другий `for` цикл всередині функції `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
len(outer_loop.find_bodies()[0].find_for_loops()) >= 1
      `)
      );
    }
});
```

Ваш `for` цикл має ітерувати над `range(n)` з `node_no` як змінною ітерації.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
outer_loop = func.find_for_loops()[0]
inner_loop = outer_loop.find_bodies()[0].find_for_loops()[0]
assert inner_loop.find_for_vars().is_equivalent("node_no")
assert inner_loop.find_for_iter().is_equivalent("range(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1

--fcc-editable-region--
```
