---
id: 68760464eb9c2e79cc913fac
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

На цьому кроці ви додасте цикл, який виконуватиметься один раз для кожного вузла в графі. Цей цикл дозволить алгоритму оновлювати відстані та шляхи протягом кількох проходів.

Створіть `for` цикл, який виконається `n` разів. Використовуйте `_` як змінну циклу, оскільки вам не потрібно використовувати значення ітерації.

Всередині циклу потрібно підготуватися до вибору наступного вузла для обробки, створивши дві змінні:

- одну для збереження найменшої відстані, знайденої на цей момент у поточній ітерації  
- і іншу для збереження індексу вузла, який має цю найменшу відстань.

Створіть змінні `min_distance` і `current` та встановіть їх відповідно в `INF` і `-1`.

# --hints--

У вас має бути `for` цикл всередині функції `shortest_path`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0]`
        )
      );
    }
});
```

Ваш `for` цикл має ітерувати над `range(n)`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_iter().is_equivalent("range(n)")`
        )
      );
    }
});
```

Ваш цикл for має використовувати `_` як змінну ітерації.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_for_vars().is_equivalent("_")`
        )
      );
    }
});
```

Ви маєте створити змінну з ім’ям `min_distance` і встановити її в `INF` всередині циклу for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("min_distance = INF")`
        )
      );
    }
});
```

Ви маєте створити змінну з ім’ям `current` і встановити її в `-1` всередині циклу for.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].has_stmt("current = -1")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
--fcc-editable-region--

--fcc-editable-region--
```
