---
id: 68760c4c2648e7832c5eef16
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Після циклу, який знаходить найближчий невідвіданий вузол, потрібно перевірити, чи дійсно знайдено валідний вузол.

Якщо такого вузла немає, це означає, що решта вузлів недосяжні зі стартового вузла, і алгоритм має зупинитися раніше.

На тому ж рівні, що й вкладений `for` цикл, додайте інструкцію `if`, яка перевіряє, чи `current == -1`, і виходить із циклу, якщо це правда.

# --hints--

У вас має бути інструкція `if`, яка перевіряє, чи `current` все ще `-1` після внутрішнього `for` циклу.

```js
({
    test: () => {
      runPython(`
      cond = _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
      assert cond.is_equivalent("current == -1") or cond.is_equivalent("-1 == current")
      `);
    }
});
```

Використайте інструкцію `break`, щоб вийти з циклу.

```js
({
    test: () => {
      runPython(`
      assert _Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt("break")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            if not visited[node_no] and distances[node_no] < min_distance:  
                min_distance = distances[node_no]  
                current = node_no

--fcc-editable-region--

--fcc-editable-region--
```
