---
id: 6876089f7a500b7c782e5b83
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Вам потрібно вирішити, чи є поточний вузол кращим вибором, ніж той, який ви вже знайшли (якщо такий є). Для цього додайте умовну інструкцію всередину циклу.

Умова має виконувати дві дії:

- Перевірити, чи вузол ще не був відвіданий.
- Порівняти відому відстань до поточного `min_distance`.

Всередині внутрішнього циклу `for` додайте інструкцію `if`, яка перевіряє, чи `node_no` не був відвіданий **і** чи `distances[node_no]` менший за `min_distance`.

Додайте `pass` як заповнювач всередині блоку умовної інструкції.

# --hints--

Ви маєте додати інструкцію `if` всередину вашого внутрішнього циклу `for`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0]`
        )
      );
    }
});
```

Ваша інструкція `if` має перевіряти, чи вузол не відвідувався (`visited[node_no]` є хибним) і чи `distances[node_no]` менший за `min_distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("not visited[node_no] and distances[node_no] < min_distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

--fcc-editable-region--
    for _ in range(n):  
        min_distance = INF  
        current = -1  
        for node_no in range(n):  
            pass

--fcc-editable-region--
```
