---
id: 68777816e95d4936f129819f
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

Перед тим як спробувати оновити відстань до сусіда, потрібно перевірити, що сусід доступний і ще не відвідувався. Потім ви обчислите, якою буде загальна відстань, щоб дістатися до цього сусіда через поточний вузол.

Всередині циклу `for` додайте інструкцію `if`, яка перевіряє:

- Що `distance` не дорівнює `INF` (це означає, що між вузлами є ребро)
- Що сусідній вузол `node_no` ще не відвідувався

Всередині блоку `if` створіть змінну з ім’ям `new_distance` і призначте їй суму `distances[current]` (найкоротша відстань до поточного вузла) та `distance` (відстань від поточного вузла до сусіда).

# --hints--

Ви повинні мати інструкцію `if`, яка перевіряє, чи `distance` не дорівнює `INF` і чи сусідній вузол ще не відвідувався.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Ви повинні створити змінну з ім’ям `new_distance` всередині блоку `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Ви повинні призначити `new_distance` суму `distances[current]` і `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
