---
id: 68e4c0ff9edcc92771d7549a
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Зараз список `distances` містить лише один елемент. Однак вам потрібно одне значення відстані для кожного вузла в графі.

У Python можна помножити список на ціле число, щоб повторити його. Наприклад, `[0] * 3` створює `[0, 0, 0]`.

Оновіть список `distances`, щоб він містив `n` копій `INF`, помноживши `[INF]` на `n`.

# --hints--

Ви маєте оновити `distances`, щоб він містив `n` копій `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
