---
id: 66db529d37ad966480ebb633
title: 建置一個午餐選擇器程式
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

在本實驗中，你將建置一個有助於管理午餐選項的程式。你將操作一個午餐陣列，為陣列添加和移除條款，並隨機選擇一個午餐選項。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該創建一個 `lunches` 變數並指定一個空的陣列，該陣列將用來儲存午餐項目。

2. 你應該創建一個函式（程式）`addLunchToEnd`，該函式（程式）以陣列作為第一個引數，以字串作為第二個引數。該函式（程式）應該：
    - 將字串添加到你的陣列末端。
    - 將字串 `[Lunch Item] added to the end of the lunch menu.` 記錄到主控台，其中 `[Lunch Item]` 是傳遞給函式的字串。
    - 傳回更新後的陣列。

3. 你應該創建一個函式（程式）`addLunchToStart`，該函式以陣列作為第一個引數，以字串作為第二個引數。該函式應該：
    - 將字串添加到你的陣列開頭。
    - 將字串 `[Lunch Item] added to the start of the lunch menu.` 記錄到主控台，其中 `[Lunch Item]` 是傳遞給函式的字串。
    - 傳回更新後的陣列。

4. 你應該創建一個函式（程式）`removeLastLunch`，該函式（程式）以一個陣列作為其引數。該函式（程式）應該：
   - 從陣列中移除最後一個元素。
   - 如果移除成功，請將字串 `[Lunch Item] removed from the end of the lunch menu.` 記錄到主控台，其中 `[Lunch Item]` 是從陣列中移除的項目。
   - 如果陣列為空，請將字串 `"No lunches to remove."` 記錄到主控台。
   - 傳回更新後的陣列。

5. 你應該創建一個函式（程式）`removeFirstLunch`，該函式以一個陣列作為其引數。該函式應該：
   - 從陣列中移除第一個元素。
   - 如果移除成功，將字串 `[Lunch Item] removed from the start of the lunch menu.` 記錄到主控台，其中 `[Lunch Item]` 是從陣列中移除的項目。
   - 如果陣列為空，請將字串 `"No lunches to remove."` 記錄到主控台。
   - 傳回更新後的陣列。

6. 你應該創建一個函式（程式）`getRandomLunch`，該函式以一個陣列作為其引數。該函式應該：
   - 從陣列中選擇一個隨機元素。
   - 如果成功，將字串 `Randomly selected lunch: [Lunch Item]` 記錄到主控台，其中 `[Lunch Item]` 是陣列中的隨機項目。
   - 如果陣列為空，請將字串 `"No lunches available."` 記錄到主控台。

7. 你應該創建一個函式（程式）`showLunchMenu`，該函式（程式）以一個陣列作為其引數，並且：
   - 如果陣列中有元素，則會在主控台記錄字串 `Menu items: [Lunch Item], [Lunch Item]...`，其中每個 `[Lunch item]` 都是陣列中的一個元素，且依序排列。
   - 如果陣列為空，則將字串 `"The menu is empty."` 記錄到主控台。

# --hints--

你應該宣告一個變數 `lunches` 並將它指定為一個空的陣列來儲存午餐項目。

```js
assert.isArray(lunches);
```

你應該定義一個函式（程式）`addLunchToEnd`。

```js
assert.isFunction(addLunchToEnd);
```

函式（程式）`addLunchToEnd` 應該有兩個參數。

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` 應該會在主控台記錄字串 `"Tacos added to the end of the lunch menu."`。

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` 應該傳回 `["Pizza", "Tacos", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

你應該定義一個函式（程式）`addLunchToStart`。

```js
assert.isFunction(addLunchToStart);
```

函式（程式）`addLunchToStart` 應該有兩個參數。

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` 應該會在主控台記錄字串 `"Sushi added to the start of the lunch menu."`。

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` 應該傳回 `["Pizza", "Burger", "Sushi"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

你應該定義一個函式（程式）`removeLastLunch`。

```js
assert.isFunction(removeLastLunch)
```

函式（程式）`removeLastLunch` 應該有一個參數。

```js
assert.lengthOf(removeLastLunch, 1);
```

當輸入陣列為空時，函式 `removeLastLunch` 應該將字串 `"No lunches to remove."` 記錄到主控台。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` 應該會在主控台記錄字串 `"Toast removed from the end of the lunch menu."`。

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` 應該傳回 `["Sushi", "Pizza"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

你應該定義一個函式（程式）`removeFirstLunch`。

```js
assert.isFunction(removeFirstLunch);
```

函式（程式）`removeFirstLunch` 應該只有一個參數。

```js
assert.lengthOf(removeFirstLunch, 1);
```

當輸入的陣列為空時，函式（程式）`removeFirstLunch` 應該在主控台記錄字串 `"No lunches to remove."`。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` 應該會在主控台記錄字串 `"Salad removed from the start of the lunch menu."`。

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` 應該傳回 `["Pizza", "Burger"]`。

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`、`addLunchToStart`、`removeLastLunch` 和 `removeFirstLunch` 應該在更新後傳回相同的作為引數傳入的陣列。

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

你應該定義一個函式（程式）`getRandomLunch`。

```js
assert.isFunction(getRandomLunch);
```

函式（程式）`getRandomLunch` 應該只有一個參數。

```js
assert.lengthOf(getRandomLunch, 1);
```

當輸入陣列為空時，函式 `getRandomLunch` 應該將字串 `"No lunches available."` 記錄到主控台。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

當輸入陣列不為空時，函式 `getRandomLunch` 應該在主控台記錄格式為 `Randomly selected lunch: [Lunch Item]` 的字串。

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

`getRandomLunch` 函式不應該更改傳入的陣列引數。

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

你應該定義一個函式（程式）`showLunchMenu`。

```js
assert.isFunction(showLunchMenu);
```

函式（程式）`showLunchMenu` 應該只有一個參數。

```js
assert.lengthOf(showLunchMenu, 1);
```

當輸入陣列為空時，函式（程式）`showLunchMenu` 應該將字串 `"The menu is empty."` 記錄到主控台。

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` 應該會在主控台記錄 `"Menu items: Greens, Corns, Beans"`。

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` 應該會在主控台記錄 `"Menu items: Pizza, Burger, Fries, Salad"`。

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
