---
id: 669e81368e52b3a5c35a2dc5
title: Crie uma página de compilação de vídeos
challengeType: 14
dashedName: build-a-video-compilation-page
---

# --description--


**Objetivo:** atender às histórias de usuário abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de usuário:**

1. Você deve ter apenas um elemento `main`.
1. Você deve ter um elemento `h1` com o tema da página.
1. Você deve ter um parágrafo introduzindo o tema da página abaixo do elemento `h1`.
1. Você deve ter três elementos `section`, cada um contendo um elemento `h2`, um parágrafo e um elemento `iframe`.
1. Os três elementos `iframe` devem ter um atributo `src` definido para um vídeo válido.
1. Cada elemento `iframe` também deve ter um atributo `title` para descrever o conteúdo incorporado, além de um atributo `height` e um atributo `width` para definir o elemento em um tamanho adequado.

# --hints--

Você deve ter um elemento `main` dentro do elemento `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

O elemento `main` deve ser o único filho do elemento `body`.

```js
assert.equal(document.querySelector('body').children.length, 2);
```

Você deve ter um elemento `h1` com o tema da página dentro do elemento `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Você deve ter um parágrafo introduzindo o tema da página abaixo do elemento `h1`.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Você deve ter três elementos `section` abaixo do primeiro elemento `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Cada elemento `section` deve conter um elemento `h2` para dar um título a essa seção como primeiro filho.

```js
const sections = $('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Cada elemento `section` deve conter um elemento `p` para introduzir o conteúdo do vídeo como segundo filho.

```js
const sections = $('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Cada elemento `section` deve conter um elemento `iframe` como terceiro filho.

```js
const sections = $('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

O primeiro elemento `iframe` deve ter um atributo `src` definido para um vídeo válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

O primeiro elemento `iframe` deve ter um atributo `title` para descrever o conteúdo incorporado.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

O primeiro elemento `iframe` deve ter um atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

O primeiro elemento `iframe` deve ter um atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

O segundo elemento `iframe` deve ter um atributo `src` definido para um vídeo válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

O segundo elemento `iframe` deve ter um atributo `title` para descrever o conteúdo incorporado.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

O segundo elemento `iframe` deve ter um atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

O segundo elemento `iframe` deve ter um atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

O terceiro elemento `iframe` deve ter um atributo `src` definido para um vídeo válido.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

O terceiro elemento `iframe` deve ter um atributo `title` para descrever o conteúdo incorporado.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

O terceiro elemento `iframe` deve ter um atributo `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

O terceiro elemento `iframe` deve ter um atributo `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front End Web Development</h1>
        <p>Front End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front End
            evelopers
            ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
