---
id: 6ac0ae563f6de70930f47ef7
title: 第 4 步
challengeType: 20
dashedName: step-4
---

# --description--

编写代码时，你经常会想要重用或引用代码中已有的值。变量可以帮你做到这一点。

要为变量赋值，写下变量名，后面跟上赋值操作符（`=`）和你想赋给该变量的值：

```py
sentence = 'I am learning Python'
```

转到高亮的那一行，创建一个名为 `sentence` 的变量，并赋值字串 `I am learning Python`。

# --hints--

你的代码中应该有 `sentence = 'I am learning Python'`。确保从行首开始写。

```js
({ test: () => runPython(`
import re
target_var = _Node(_code).find_variable("sentence")
pattern = r"""sentence\\s*=\\s*(['"])\\s*i\\s+am\\s+learning\\s+python\\s*\\1\\s*"""
assert re.fullmatch(pattern, str(target_var).strip(), flags=re.IGNORECASE)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')
--fcc-editable-region--

--fcc-editable-region--
```
