---
id: 6ac09f8420b4b67e075d2cd2
title: 第 1 步
challengeType: 20
dashedName: step-1
---

# --description--

在本次工作坊中，你将创建你的第一个 Python 程序。

这是你的代码编辑器，你将在这里编写 Python 代码。第 1 行已高亮。去那里并输入：

```py
'Hello'
```

确保从行首开始输入，并且不要忘记引号。

完成后，点击“查看你的代码”按钮，看看是否正确。

# --hints--

你应该输入 `'Hello'`。确保从行首开始输入，并且不要忘记引号。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
