---
id: 6ac0a4463662d9738c393e21
title: 第 2 步
challengeType: 20
dashedName: step-2
---

# --description--

在上一步中，你创建了一个**字串**：由单引号或双引号包裹的一系列字符。在编程中，字串用于表现文本。

**终端**是你可以看到代码输出的地方。

为了能够看到输出，将你现有的代码行改成：

```py
print('Hello')
```

确保从行首开始。

# --hints--

你的代码中应该有 `print('Hello')`。确保从行首开始。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
