---
id: 6ac0b35f3b29a27f7e207295
title: 第 5 步
challengeType: 20
dashedName: step-5
---

# --description--

在上一步中，你创建了变量 `sentence`。现在通过在高亮行输入 `print(sentence)` 将其值打印到终端。

这样，你就完成了你的第一个 Python 程序。

# --hints--

你的代码中应该有 `print(sentence)`。确保从行首开始输入。

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
