---
id: 6ac0aaba914d81b32567c076
title: 第 3 步
challengeType: 20
dashedName: step-3
---

# --description--

`print` 是一个内置函数，你可以用它向终端发送输出。正如你所见，文本 `Hello` 现在已经显示在终端上。

现在，转到高亮的行并输入，将文本 `freeCodeCamp` 打印到终端：

```py
print('freeCodeCamp')
```

确保从行首开始输入。

# --hints--

你的代码中应该有 `print('freeCodeCamp')`。确保从行首开始输入。

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
