---
id: 646c47867800472a4ed5d2ea
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza na muundo wa msingi wa HTML. Ongeza `DOCTYPE` rejea ya `html` na kipengele cha `html` chenye sifa ya kipengele `lang` iliyowekwa kuwa `en`. Pia, ongeza kipengele cha `head` na `body` ndani ya kipengele cha `html`.

# --hints--

Unapaswa kuwa na tamko la `DOCTYPE` la `html`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Unapaswa kuwa na lebo ya ufunguzi ya `html` yenye sifa ya kipengele `lang` ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Unapaswa kuwa na lebo ya kufunga ya `html`.

```js
assert(code.match(/<\/html>/i));
```

Unapaswa kuwa na lebo ya ufunguzi ya `head`.

```js
assert(code.match(/<head>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `head`.

```js
assert(code.match(/<\/head>/i));
```

Unapaswa kuwa na lebo ya ufunguzi ya `body`.

```js
assert(code.match(/<body>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `body`.

```js
assert(code.match(/<\/body>/i));
```

Kipengele chako cha `body` kinapaswa kuja baada ya kipengele cha `head`.

```js
assert(code.match(/<head>\s*<\/head>\s*<body>\s*<\/body>/i));
```

Vipengele vyako vya `head` na `body` vinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<head>\s*<\/head>\s*<body>\s*<\/body>\s*<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
