---
id: 646ce9d790d2a44de5f99e04
title: Hatua 21
challengeType: 0
dashedName: step-21
---

# --description--

Unda vipengele viwili vya `div`, cha kwanza ndani ya kipengele cha `.cat-left-ear` chenye darasa la vitu la `cat-left-inner-ear`, na cha pili ndani ya kipengele cha `.cat-right-ear` chenye darasa la vitu la `cat-right-inner-ear`.

# --hints--

Haupasi kubadilisha kipengele cha `div` kilichopo chenye darasa la vitu la `cat-left-ear`.

```js
assert(document.querySelectorAll('div.cat-left-ear').length === 1);
```

Haupasi kubadilisha kipengele cha `div` kilichopo chenye darasa la vitu la `cat-right-ear`.

```js
assert(document.querySelectorAll('div.cat-right-ear').length === 1);
```

Unapaswa kuwa na kipengele kimoja cha `div` ndani ya kipengele chako cha `.cat-left-ear`.

```js
assert(document.querySelectorAll('.cat-left-ear div').length === 1);
```

Unapaswa kuwa na kipengele kimoja cha `div` ndani ya kipengele chako cha `.cat-right-ear`.

```js
assert(document.querySelectorAll('.cat-right-ear div').length === 1);
```

Kipengele kipya cha `div` ndani ya `.cat-left-ear` kinapaswa kuwa na darasa la vitu la `cat-left-inner-ear`.

```js
assert(document.querySelectorAll('.cat-left-ear div')[0]?.classList.contains('cat-left-inner-ear'));
```

Kipengele kipya cha `div` ndani ya `.cat-right-ear` kinapaswa kuwa na darasa la vitu la `cat-right-inner-ear`.

```js
assert(document.querySelectorAll('.cat-right-ear div')[0]?.classList.contains('cat-right-inner-ear'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          --fcc-editable-region--
          <div class="cat-left-ear">

          </div>
          <div class="cat-right-ear">
            
          </div>
          --fcc-editable-region--
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
