---
id: 646c5d7057c45f432fcdd46c
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

Kutumia kichaguzi cha darasa la vitu, mpe kipengele cha `.cat-head` upana wa `205px` na urefu wa `180px`. Pia, mpe mpaka wa `1px solid #000` na `border-radius` ya `46%`. 

# --hints--

Unapaswa kuwa na kichaguzi cha `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na `width` iliyowekwa kuwa `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na `height` iliyowekwa kuwa `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na `border` iliyowekwa kuwa `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na `border-radius` iliyowekwa kuwa `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
