---
id: 6476fb7cbaafa36d65e9cf35
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

`fixed` ni thamani ya sifa ya `position` inayokuwezesha kuweka kipengele kikiwa kimefungwa kwenye ukurasa bila kujali mtumizi anaposogeza ukurasa.

Utahitaji kufanya alama za uandishi zaidi ili kuona jinsi nafasi ya `fixed` inavyofanya kazi. Katika HTML yako, tengeneza kipengele cha `div` chenye darasa la `box`.

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert(document.querySelectorAll('div').length === 2);
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la `box`.

```js
assert(document.querySelectorAll('div')[1]?.getAttribute('class') === 'box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  top: 300px;
  left: 400px;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
