---
id: 6476f5c17f99146071ee884c
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Uliweza kuona kwamba hakuna kinachotokea. Kipengele cha `.cat-head` hakikuhamishwa licha ya kuweka `top` na `left` za `100px` kila moja. Lakini hali si hivyo kwa upangaji wa `relative`.

Unapotumia thamani ya `relative`, kipengele bado kinawekwa kulingana na mtiririko wa kawaida wa hati, lakini thamani za `top`, `left`, `bottom`, na `right` zinakuwa hai.

Badala ya `static`, mpe `.cat-head` yako nafasi ya `relative`, na acha sifa za `top` na `left` kama zilivyo.

# --hints--

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `position` iliyowekwa kuwa `relative`. Hakikisha unaongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'relative')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `top` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Kichaguzi chako cha `.cat-head` kinapaswa kuwa na sifa ya `left` iliyowekwa kuwa `100px`. Hakikisha unaongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: static;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
