---
id: 647d855ac12cd436059acd39
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Tumia kichaguzi cha ulimwengu wote kuongeza `box-sizing: border-box;` kwenye CSS yako. Hii inahakikisha vipengele vinajumuisha nafasi ya ndani na mpaka katika upana na urefu vilivyoainishwa.

# --hints--

Unapaswa kuwa na kichaguzi cha ulimwengu wote (`*`).

```js
// assert.match(code, /\*\s*{\s*.*\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*'))
```

Kichaguzi chako cha ulimwengu wote kinapaswa kuwa na sifa ya `box-sizing` iliyowekwa kuwa `border-box`.

```js
// assert.match(code, /\*\s*{\s*box-sizing:\s*border-box;?\s*}/)
assert(new __helpers.CSSHelp(document).getStyle('*')?.boxSizing === 'border-box')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body></body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
