---
id: 6676a8b01e56ec1a1e07c202
title: Hatua 23
challengeType: 0
dashedName: step-23
---

# --description--

Tambua kuwa sasa una mraba mweupe wenye mipaka nene. Mipaka hiyo ina pembe za mwinuko ambazo zinaweza kutumika kuunda mviringo.

Ndani ya kichaguzi cha darasa hilo hilo rekebisha `height` na `width` kuwa `0`. Weka mipaka ya kushoto, kulia na juu kuwa `transparent`.

Ondoa sifa ya `background-color`, na utauona mviringo.

# --hints--


Unapaswa kuwa na kichaguzi cha `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `height` iliyowekwa kuwa `0`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `0`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-left` iliyowekwa kuwa `35px solid transparent`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-right` iliyowekwa kuwa `35px solid transparent`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-top` iliyowekwa kuwa `35px solid transparent`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


Kichaguzi chako cha `.cat-right-ear` kinapaswa kuwa na sifa ya `border-bottom` iliyowekwa kuwa `35px solid red`. Usisahau kuongeza nukta ya nusu.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

Kichaguzi chako cha `.cat-right-ear` hakipaswi kuwa na sifa ya `background-color`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
