---
id: 64dc97005bc7943e2978df0a
title: Passaggio 11
challengeType: 20
dashedName: step-11
---

# --description--

Ora puoi rimuovere la tua chiamata `print()`. Poi, all'interno della funzione `move()`, elimina la parola chiave `pass` e stampa il contenuto del tuo dizionario `rods`.

# --hints--

Non dovresti avere `print(number_of_moves)` nel tuo codice.

```js
({ test: () => assert.isFalse(/print\s*\(\s*number_of_moves\s*\)/.test(code)) })
```

Dovresti rimuovere la parola chiave `pass` e stampare il dizionario `rods`.

```js
({ test: () => assert.match(code, /def\s+move\s*\(\s*\)\s*:\s+print\s*\(\s*rods\s*\)(?!\s*pass)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
print(number_of_moves)
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
--fcc-editable-region--
```
