---
id: 64dc922df2919e3c38cead61
title: Passaggio 8
challengeType: 20
dashedName: step-8
---

# --description--

All'inizio del tuo codice, dichiara una variabile chiamata `NUMBER_OF_DISKS` per archiviare il numero di dischi e assegnale il valore `3`. Poi, sostituisci il primo argomento passato alla funzione `range()` con la tua nuova variabile.

# --hints--

Dovresti dichiarare una variabile chiamata `NUMBER_OF_DISKS`.

```js
({test: () => assert(__userGlobals.has('NUMBER_OF_DISKS')) })
```

La variabile `NUMBER_OF_DISKS` dovrebbe avere il valore `3`.

```js
({test: () => assert.equal(__userGlobals.get('NUMBER_OF_DISKS'), 3) })
```

Dovresti cambiare il valore del primo argomento della funzione `range()` nella variabile `NUMBER_OF_DISKS`.

```js
({test: () => assert.match(code, /('|")A\1\s*:\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

def move():
    pass
```
