---
id: 64dc92a9718fb53ca3d1bad7
title: Passaggio 9
challengeType: 20
dashedName: step-9
---

# --description--

Il puzzle della torre di Hanoi può essere risolto in 2<sup>n</sup> - 1 mosse, dove n è il numero di dischi. Dichiara una variabile chiamata `number_of_moves` e assegna a questa variabile il numero totale di mosse.

L'operatore di potenza in Python è `**`.

# --hints--

Dovresti dichiarare una variabile chiamata `number_of_moves`.

```js
({ test: () => assert(__userGlobals.has('number_of_moves')) })
```

Il valore di `number_of_moves` dovrebbe essere l'espressione per calcolare il numero di mosse.

```js
({ test: () => {
    assert(runPython(`  
    number_of_moves == 2**NUMBER_OF_DISKS -1      
    `))
} })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 3

--fcc-editable-region--
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    pass
```
