---
id: 64dc985e6720a23edac63b51
title: Passaggio 17
challengeType: 20
dashedName: step-17
---

# --description--

I movimenti consentiti del disco tra i pali seguono un pattern ripetitivo che si verifica ogni tre mosse. Per esempio, i movimenti tra il palo `A` e il palo `C` sono permessi alla prima, alla quarta e alla settima mossa, dove il resto della divisione tra il numero della mossa e 3 è uguale a 1.

All'interno del loop `for` creato in precedenza, sostituisci la chiamata `print()` esistente con un'istruzione `if` che viene attivata quando `(i + 1) % 3 == 1`. All'interno di questa istruzione `if`, stampa `f'Move {i + 1} allowed between {source} and {target}'` usando una f-string. Nota che `i + 1` è il numero della mossa dato che `i` è zero durante la prima iterazione.

# --hints--

Dovresti avere un'istruzione `if` che viene eseguita quando `(i + 1) % 3 == 1`. 

```js
({test: () => assert.match(code, /if\s+\(\s*i\s*\+\s*1\s*\)\s*%\s*3\s*==\s*1\s*:/)})
```

Dovresti stampare la stringa `f'Move {i + 1} allowed between {source} and {target}'` usando una f-string.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*target\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        print(i)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
