---
id: 64dc8ea01436383a88256d7f
title: Passaggio 3
challengeType: 20
dashedName: step-3
---

# --description--

Il puzzle inizia con i dischi impilati sulla prima asta, in ordine decrescente di dimensione, con il disco più piccolo in cima e quello più grande in fondo. Devi popolare la tua prima lista con numeri che rappresentano le varie dimensioni dei dischi.

Invece di aggiungere manualmente gli elementi alla prima lista, genera una sequenza di numeri che contano all’indietro da `3` a `1` usando la funzione `range()` e assegnala a `rods['A']`. Qui, `3` rappresenta il disco più grande in fondo alla pila e `1` rappresenta il disco più piccolo in cima alla pila.

La sintassi è `range(x, y, h)`, dove `x` è l’intero iniziale (incluso), `y` è l’ultimo intero (non incluso) e `h` è la differenza tra un numero e il successivo nella sequenza.

# --hints--

Dovresti usare la funzione `range()` per assegnare una sequenza di numeri a `rods['A']`. La sintassi per chiamare la funzione `range()` è `range(x, y, h)`.

```js
({ test: () => assert(runPython(`
  type(rods['A']) is range
  `))
})
```

Dovresti usare la funzione `range()` per assegnare la sequenza di numeri da `3` a `1` a `rods['A']`.

```js
({ test: () => assert(runPython(`
 rods['A'] == range(3, 0, -1)
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': [],
    'B': [],
    'C': []
}
--fcc-editable-region--
```
