---
id: 64dc976bf864693e668d67e8
title: Passaggio 12
challengeType: 20
dashedName: step-12
---

# --description--

Ora chiama la tua funzione e guarda l'output sul terminale.

# --hints--

Dovresti chiamare la tua funzione `move()`.

```js
({test: () => assert.match(code, /^move\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
def move():
    print(rods)
--fcc-editable-region--

--fcc-editable-region--
```
