---
id: 64dcd3d61c448e2676501f43
title: Passaggio 20
challengeType: 20
dashedName: step-20
---

# --description--

Quando il resto della divisione del numero della mossa per 3 è uguale a 2, il movimento è consentito tra `'A'` e `'B'` (la bacchetta di partenza e quella ausiliaria).

Aggiungi un'istruzione `elif` per questo caso. Poi, stampa la stringa appropriata se la condizione è soddisfatta.

# --hints--

Dovresti avere un'istruzione `elif` da eseguire quando `remainder == 2`.

```js
({test: () => assert.match(code, /elif\s+remainder\s*==\s*2:/)})
```

Dovresti stampare la stringa `f'Move {i + 1} allowed between {source} and {auxiliary}'`.

```js
({test: () => assert.match(code, /print\s*\(\s*f('|")Move\s\{\s*i\s*\+\s*1\s*\}\sallowed\sbetween\s\{\s*source\s*\}\sand\s\{\s*auxiliary\s*\s*\}\1\s*\)/)})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
