---
id: 64dc97ab25730b3e9de86ea5
title: Passaggio 16
challengeType: 20
dashedName: step-16
---

# --description--

Alla fine di questo progetto, creerai una soluzione ricorsiva per il puzzle della torre di Hanoi, ma ora esplorerai un approccio iterativo a questo problema.

Inizia aggiungendo un loop (noun) `for` alla tua funzione che itera su `number_of_moves` e stampa il numero dell'iterazione corrente.

# --hints--

Dovresti scrivere un loop (noun) `for` per iterare sul numero di mosse. Usa la funzione `range()` per farlo.

```js
({ test: () => {
    const rgs = [
        /for\s+\w+\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:/,
        /for\s+\w+\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

Dovresti stampare il numero della mossa corrente a ogni iterazione.

```js
({ test: () => {
    const rgs = [
        /for\s+(\w+)\s+in\s+range\s*\(\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*\)\s*:\s+print\s*\(\s*\1\s*\)/,
        /for\s+(\w+)\s+in\s+range\s*\(\s*0\s*,\s*number_of_moves\s*,\s*1\s*\)\s*:\s+print\s*\(\s*\1\s*\)/
    ]
    const loop = rgs.some(r => code.match(r));
    assert.isTrue(loop);
  }
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
--fcc-editable-region--

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
