---
id: 64dc90949ff85e3b37be40d0
title: Passaggio 5
challengeType: 20
dashedName: step-5
---

# --description--

La funzione `range()` restituisce una sequenza immutabile di numeri. Come puoi vedere, il tipo di dato di `rods['A']` è `range`, ma vuoi che sia una lista.

Passa la tua chiamata `range()` alla funzione `list()` per farlo.

# --hints--

Dovresti passare la tua chiamata `range()` alla funzione `list()`.

```js
({ test: () => assert(runPython(`
  rods['A'] == list(range(3, 0, -1))
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
print(type(rods['A']))
--fcc-editable-region--
```
