---
id: 64dc916138dbe03bdfe6cec6
title: Passaggio 7
challengeType: 20
dashedName: step-7
---

# --description--

L'obiettivo della torre di Hanoi è spostare tutti i dischi sull'ultimo palo. Per farlo, devi seguire tre semplici regole:

1. Puoi spostare solo i dischi più in alto
2. Puoi spostare solo un disco alla volta
3. Non puoi mettere dischi più grandi sopra quelli più piccoli

Sotto il tuo codice esistente, dichiara una funzione vuota chiamata `move`. In seguito userai quella funzione per spostare i dischi tra i pali. Per ora, per evitare errori, usa la parola chiave `pass` all'interno del corpo della funzione.

# --hints--

Dovresti dichiarare una funzione vuota chiamata `move`. Ricorda di usare la parola chiave `pass` all'interno del corpo della funzione con la corretta indentazione.

```js
({ test: () => {
		assert(runPython(`
        import inspect
        inspect.isfunction(move)
    `))
	} 
})
```

# --seed--

## --seed-contents--

```py
rods = {
    'A': list(range(3, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--
```
