---
id: 64dc8b1faf9c60304ca8a9b6
title: Passaggio 1
challengeType: 20
dashedName: step-1
---

# --description--

In questo progetto risolverai il rompicapo matematico noto come torre di Hanoi. Il rompicapo consiste in tre aste e un numero di dischi di diametri diversi.

L’obiettivo di questo rompicapo è spostare i dischi dalla prima asta alla terza, seguendo regole specifiche che vietano di posizionare un disco più grande sopra uno più piccolo.

Inizia creando un dizionario vuoto chiamato `rods` per rappresentare le aste.

# --hints--

Dovresti avere una variabile chiamata `rods`.

```js
({ test: () => assert(__userGlobals.has('rods')) })
```

La tua variabile `rods` dovrebbe essere un dizionario vuoto.

```js
({ test: () => assert(runPython(`
    rods == {}
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
