---
id: 64dc9004f658183afa85f4f1
title: Passaggio 4
challengeType: 20
dashedName: step-4
---

# --description--

Ora controlla il tipo di dato della tua chiave `'A'` passandola alla funzione `type()` e stampa il risultato sul terminale.

# --hints--

Dovresti passare la tua chiave `'A'` alla funzione `type()`.

```js
({ test: () => assert.match(code, /type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)/) })
```

Dovresti stampare il tipo di `rods['A']` usando la funzione `print()`.

```js
({ test: () => assert.match(code, /print\s*\(\s*type\s*\(\s*rods\s*\[\s*('|")\s*A\s*\1\s*\]\s*\)\s*\)/) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {
    'A': range(3, 0, -1),
    'B': [],
    'C': []
}
--fcc-editable-region--
```
