---
id: 64ddd280d69fa818ca782d50
title: Passaggio 29
challengeType: 20
dashedName: step-29
---

# --description--

Come puoi vedere, il disco 1 si muove avanti e indietro ogni tre mosse. Questo succede perché devi ancora gestire i movimenti tra gli altri pali.  
Invece di ripetere lo stesso codice che hai scritto nei passaggi precedenti cambiando i pali, sarebbe meglio spostare quel codice all'interno di una funzione da chiamare in ogni istruzione condizionale. Dichiara una funzione vuota chiamata `make_allowed_move()` e non dimenticare la parola chiave `pass`.

# --hints--

Dovresti dichiarare una funzione vuota chiamata `make_allowed_move`. Ricorda di usare la parola chiave `pass` all'interno del corpo della funzione con la corretta indentazione.

```js
({ test: () => assert(runPython(`
      import inspect
      inspect.isfunction(make_allowed_move)
  `)) 
})
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 3
number_of_moves = 2**NUMBER_OF_DISKS - 1
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}
--fcc-editable-region--

--fcc-editable-region--

def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods)
    for i in range(number_of_moves):
        remainder = (i + 1) % 3
        if remainder == 1:
            print(f'Move {i + 1} allowed between {source} and {target}')
            forward = False
            if not rods[target]:
                forward = True
            elif rods[source] and rods[source][-1] < rods[target][-1]:
                forward = True              
            if forward:
                print(f'Moving disk {rods[source][-1]} from {source} to {target}')
                rods[target].append(rods[source].pop())
            else:
                print(f'Moving disk {rods[target][-1]} from {target} to {source}')
                rods[source].append(rods[target].pop())
            
            # display our progress
            print(rods)
        elif remainder == 2:
            print(f'Move {i + 1} allowed between {source} and {auxiliary}')
        elif remainder == 0:
            print(f'Move {i + 1} allowed between {auxiliary} and {target}')

# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
