---
id: 64de1469fccce61940f498c4
title: Passaggio 2
challengeType: 20
dashedName: step-2
---

# --description--

Il dizionario `rods` rappresenterà le tre aste con i loro dischi. Dagli come chiavi le stringhe `'A'`, `'B'` e `'C'` e imposta ciascuna di esse su una lista vuota.

# --hints--

Il tuo dizionario `rods` dovrebbe avere una chiave `'A'`.

```js
({ test: () => assert(runPython(`'A' in rods`)) })
```

`rods['A']` dovrebbe essere una lista vuota.

```js
({ test: () => assert(runPython(`
    rods['A'] == []
  `))
})
```

Il tuo dizionario `rods` dovrebbe avere una chiave `'B'`.

```js
({ test: () => assert(runPython(`'B' in rods`)) })
```

`rods['B']` dovrebbe essere una lista vuota.

```js
({ test: () => assert(runPython(`
    rods['B'] == []
  `))
})
```

Il tuo dizionario `rods` dovrebbe avere una chiave `'C'`.

```js
({ test: () => assert(runPython(`'C' in rods`)) })
```

`rods['C']` dovrebbe essere una lista vuota.

```js
({ test: () => assert(runPython(`
    rods['C'] == []
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
rods = {}
--fcc-editable-region--
```
