---
id: 64df47b32b92301a815d5ef8
title: Passaggio 49
challengeType: 20
dashedName: step-49
---

# --description--

All’inizio, la chiamata ricorsiva che hai appena aggiunto si occupa del sotto-problema di spostare `n - 1` dischi sulla seconda asta.

Per questo motivo, l’argomento `target` corrisponde alla tua seconda asta, mentre l’argomento `auxiliary` è la terza asta. Tieni presente che questi continueranno a scambiarsi man mano che la ricorsione procede.

Correggi l’ordine degli argomenti scambiando `target` e `auxiliary` nella tua chiamata ricorsiva.

# --hints--

Dovresti modificare l’ordine degli argomenti nella tua chiamata `move(n - 1, source, auxiliary, target)`.

```js
({ test:  () => assert.match(code, /move\(\s*n\s*-\s*1\s*,\s*source\s*,\s*target\s*,\s*auxiliary\s*\)/) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, auxiliary, target)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
