---
id: 64df496c6a8ddf1b38db1ed6
title: Passaggio 50
challengeType: 20
dashedName: step-50
---

# --description--

In un passaggio precedente, hai scritto il codice per spostare il disco più grande del sotto-problema al palo obiettivo.

Ora, tutto ciò che devi fare è aggiungere un'altra chiamata ricorsiva per spostare i `n - 1` dischi che hai già spostato. Copia la prima chiamata ricorsiva e incollala alla fine del blocco `if`.

Nota che gli argomenti della funzione non sono nell'ordine corretto. Cerca di capire qual è l'ordine corretto.

# --hints--

Dovresti chiamare `move(n - 1,  auxiliary, source, target)` alla fine della tua funzione `move`.

```js
({ test: () => assert.match(code, /^\s{8}move\(\s*n\s*-\s*1\s*,\s*auxiliary\s*,\s*source\s*,\s*target\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
              
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
