---
id: 64df3f1011888113fbd3d81b
title: Passaggio 44
challengeType: 20
dashedName: step-44
---

# --description--

Per risolvere il puzzle con la ricorsione, la prima cosa da fare è scomporre il problema originale in sotto-problemi più piccoli.

La configurazione finale con `n` dischi impilati sul terzo palo in ordine decrescente può essere ottenuta spostando:

- `n - 1` dischi dalla sorgente al palo ausiliario
- il disco più grande dalla sorgente all'obiettivo
- e poi i `n - 1` dischi dal palo ausiliario all'obiettivo.

Quindi, la prima cosa che la funzione `move` dovrebbe fare è chiamare se stessa con `n - 1` come primo argomento. Ma se provi a farlo senza definire un caso base, otterrai un `RecursionError`. Questo succede perché la funzione continua a chiamare se stessa all'infinito.

Prima del tuo commento e della tua chiamata `print()`, aggiungi la chiamata ricorsiva della funzione con `n - 1` come primo argomento e assicurati che il corpo della funzione venga eseguito solo quando `n` è maggiore di zero. Per ora, lascia gli altri argomenti nello stesso ordine.

# --hints--

Il corpo della funzione `move` dovrebbe iniziare con un'istruzione `if` che viene attivata quando `n` è maggiore di zero.

```js
({ test: () => assert.match(code, /def\s+move\(\s*n\s*,\s*source\s*,\s*auxiliary\s*,\s*target\s*\)\s*:\s+if\s+n\s*>\s*0\s*:/) })
```

Dovresti chiamare `move(n - 1, source, auxiliary, target)` all'interno dell'istruzione `if`. Ricorda di indentare la chiamata `print()`.

```js
const transformedCode = "\n" + code.replace(/\r/g, '');
const move = __helpers.python.getDef(transformedCode, "move");
const { function_parameters, function_body } = move;

assert.match(function_parameters, /n\s*,\s*source\s*,\s*auxiliary\s*,\s*target/);
const functionBodyIndent = function_body.match(/ +/)[0];
const re = new RegExp(`${functionBodyIndent}if\\s+n\\s*>\\s*0\\s*:\\s+^${functionBodyIndent}( +)move\\(\\s*n\\s*-\\s*1\\s*,\\s*source\\s*,\\s*auxiliary\\s*,\\s*target\\s*\\).+?^${functionBodyIndent}\\1print\\s*\\(\\s*rods\\s*,\\s*("|')\\\\n\\2\\s*\\)`, "ms");
assert.match(function_body, re);
```

# --seed--

## --seed-contents--

```py
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

--fcc-editable-region--
def move(n, source, auxiliary, target):
    # display starting configuration
    print(rods, '\n')
    
--fcc-editable-region--
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
```
