---
id: 650c6082e5586f9e3acfcd3b
title: Passaggio 52
challengeType: 20
dashedName: step-52
---

# --description--

Ora, elimina il dizionario `rods` e trasforma le sue chiavi nelle variabili `A`, `B` e `C`, rispettivamente, mantenendo i loro valori. Rifattorizza il tuo codice per riflettere queste modifiche. Se vedi l'output sul terminale, hai fatto correttamente questo passaggio.

# --hints--

Dovresti trasformare `rods['A']` in una variabile chiamata `A` mantenendo il suo valore.

```js
assert.match(code, /A\s*=\s*list\s*\(\s*range\s*\(\s*NUMBER_OF_DISKS\s*,\s*0\s*,\s*-\s*1\s*\)\s*\)/);
```

Dovresti trasformare `rods['B']` in una variabile chiamata `B` mantenendo il suo valore.

```js
assert.match(code, /B\s*=\s*\[\s*\]/);
```

Dovresti trasformare `rods['C']` in una variabile chiamata `C` mantenendo il suo valore.

```js
assert.match(code, /C\s*=\s*\[\s*\]/);
```

Dovresti modificare la tua chiamata `print()` per stampare `A`, `B`, `C`, invece dell'oggetto `rods`. Mantieni il carattere di nuova riga nella chiamata `print()`.

```js
assert.match(code, /print\s*\(\s*A\s*,\s*B\s*,\s*C\s*,\s*('|")\\n\1\s*\)/);
```

Dovresti avere `target.append(source.pop())` nel tuo codice.

```js
assert.match(code, /target\.append\s*\(\s*source\.pop\s*\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
NUMBER_OF_DISKS = 4
rods = {
    'A': list(range(NUMBER_OF_DISKS, 0, -1)),
    'B': [],
    'C': []
}

def move(n, source, auxiliary, target):
    if n > 0:
        # move n - 1 disks from source to auxiliary, so they are out of the way
        move(n - 1, source, target, auxiliary)
        
        # move the nth disk from source to target
        rods[target].append(rods[source].pop())
        
        # display our progress
        print(rods, '\n')
        
        # move the n - 1 disks that we left on auxiliary onto target
        move(n - 1,  auxiliary, source, target)
              
# initiate call from source A to target C with auxiliary B
move(NUMBER_OF_DISKS, 'A', 'B', 'C')
--fcc-editable-region--
```
