---
id: 680a431692a5f863f726c438
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이 워크숍에서는 *시저 암호*를 만들 것입니다. 이것은 텍스트를 암호화하는 가장 간단한 기법 중 하나로, 평문 각 글자를 알파벳에서 고정된 위치만큼 떨어진 글자로 교체하는 방식입니다. 예를 들어, 5만큼 이동하면 `a`는 `f`로, `b`는 `g`로 교체됩니다.

이 암호를 구현하려면, 이동값이 가리키는 위치에서 시작하는 새로운 알파벳 버전을 만들어야 합니다. 이전 수업에서 배운 것처럼, 문자열 슬라이싱을 사용해 문자열의 일부를 추출할 수 있습니다:

```py
fcc = 'freeCodeCamp'
print(fcc[8:]) # Camp
```

`shifted_alphabet`라는 변수를 만들고, 슬라이싱 구문을 사용해 `alphabet`의 `shift` 인덱스부터 시작하는 부분을 할당하세요. 그런 다음 내장 `print()` 함수를 호출해 `shifted_alphabet`를 터미널에 출력하고 결과를 확인하세요.

# --hints--

`shifted_alphabet`라는 새 변수를 선언하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shifted_alphabet")`)) })

```

`alphabet[shift:]`에 `shifted_alphabet`을 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shifted_alphabet").is_equivalent("shifted_alphabet = alphabet[shift:]")`)) })
```

`print()`을 인수로 전달하여 `shifted_alphabet` 함수를 호출하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--

--fcc-editable-region--
```
