---
id: 680a7e200400089967b8b1f4
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`str.maketrans()` 메서드는 길이가 같은 두 문자열을 받아 첫 번째 문자열의 각 문자를 두 번째 문자열의 해당 문자와 매핑하는 변환 표를 반환합니다. 변환 표의 각 문자는 문자를 고유하게 식별하는 숫자인 유니코드 서수로 저장됩니다.

```py
lower_chars = 'abc'
upper_chars = 'ABC'

table = str.maketrans(lower_chars, upper_chars)
print(table) # {97: 65, 98: 66, 99: 67}
```

커리큘럼 후반부에서 `str.maketrans()`이 반환하는 구조의 유형에 대해 더 배우게 됩니다. 지금은 `alphabet`의 문자를 `shifted_alphabet`의 문자에 매핑하는 변환 표를 만들고 `translation_table`라는 변수에 할당하세요.

# --hints--

`translation_table`라는 변수를 선언하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("translation_table")`)) })
```

`str.maketrans()`과 `alphabet`을 인수로 전달하여 `shifted_alphabet`을 호출하고 결과를 `translation_table`에 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("translation_table").is_equivalent("translation_table = str.maketrans(alphabet, shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--

--fcc-editable-region--
```
