---
id: 680a5e7f3d47715196b12ca0
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

전체 시프트된 알파벳이 준비되었으니, `print(shifted_alphabet)` 호출을 제거하세요.

# --hints--

코드에서 `print(shifted_alphabet)`을 제거해야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).has_call("print(shifted_alphabet)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
print(shifted_alphabet)
--fcc-editable-region--
```
