---
id: 680a549f60f5a3fedfc0edd9
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

출력에서 볼 수 있듯이, `f`의 값이 `shift`이기 때문에 이동된 알파벳은 `5` 문자부터 시작합니다. 하지만 이제 이동된 알파벳에서 처음 다섯 글자 `a`, `b`, `c`, `d`, `e`가 빠졌으므로, 이 글자들을 이동된 알파벳의 끝에 추가해야 합니다.

`+` 연산자는 두 개 이상의 문자열을 연결하는 과정인 연결(concatenation)에 사용됩니다. 예를 들어 다음과 같습니다:

```py
greeting = 'Hello' + ' ' + 'World'
print(greeting) # Hello World

```

기존 `shifted_alphabet` 변수의 할당을 수정하세요: 슬라이싱 구문을 사용해 `alphabet`의 빠진 처음 부분을 추출하고 `alphabet[shift:]`에 연결하세요.

참고로, `sentence[start:stop]`는 `sentence`의 `start` 위치(포함)부터 `stop` 위치(제외)까지의 문자를 반환합니다.

# --hints--

`alphabet` 변수를 직접 수정하지 마세요.

```js
({ test: () => runPython(`
  assert alphabet == "abcdefghijklmnopqrstuvwxyz"
`) })
```

`alphabet[shift:]`과 `alphabet`의 빠진 처음 부분을 연결한 값을 `shifted_alphabet` 변수에 선언된 같은 줄에서 할당하세요. 슬라이싱을 멈출 위치를 지정하려면 `shift`을 사용하세요.

```js
({ test: () => runPython(`
_shifted_alpha = _Node(_code).find_variable("shifted_alphabet")
_first = _shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[:shift]")
_second =_shifted_alpha.is_equivalent("shifted_alphabet = alphabet[shift:] + alphabet[0:shift]")
assert _first or _second
`) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
--fcc-editable-region--
shifted_alphabet = alphabet[shift:]
--fcc-editable-region--
print(shifted_alphabet)
```
