---
id: 680a3d4c065557260de2ff33
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이전 수업에서 기억하듯이, Python에서는 대입 연산자(`=`) 왼쪽에 변수 이름을 쓰고 오른쪽에 할당할 값을 써서 변수를 선언합니다.

```py
variable_name = value
```

`shift`라는 변수를 만들고 새 변수에 `5` 값을 할당하세요.

# --hints--

`shift`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는 것에 주의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

`5` 값을 `shift` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
