---
id: 680a3d4c065557260de2ff33
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이번 워크숍에서는 알파벳의 글자를 이동시켜 메시지를 암호화하는 암호화 방법인 시저 암호 프로그램을 만들 것입니다.

`shift`라는 변수를 만들고 새 변수에 `5` 값을 할당하세요.

# --hints--

`shift`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는 것에 주의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("shift")`)) })
```

`5` 값을 `shift` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("shift").is_equivalent("shift = 5")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
