---
id: 680a41ef5bfdc059cfd9d4b5
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이전 수업에서 변수에 저장할 수 있는 다양한 데이터 유형에 대해 배웠습니다. 방금 정수 값을 할당했죠. 이제는 문자열을 할당해야 합니다. 문자열은 작은따옴표나 큰따옴표로 감싼 문자 수열입니다.

```py
string_1 = 'I am a string'
string_2 = "I am also a string"
```

`alphabet`라는 변수를 하나 더 선언하고, 이 변수에 `abcdefghijklmnopqrstuvwxyz` 문자열을 할당하세요.

# --hints--

`alphabet`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는 것에 주의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

`abcdefghijklmnopqrstuvwxyz` 변수에 `alphabet` 문자열을 할당해야 합니다. 문자열을 감쌀 때는 작은따옴표나 큰따옴표 중 하나를 사용하세요.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
