---
id: 680a836235461fce1ff53b3f
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`translate()` 메서드는 `str.maketrans()`에서 생성된 변환 표를 인수로 받습니다. 이 메서드는 문자열에서 호출되며, 변환 표에 따라 문자가 교체된 원본 문자열의 복사본을 반환합니다.

```py
t = str.maketrans('lk', 'br')
sentence = 'The tent gave in to the leaks.'

print(sentence.translate(t))
# Output: The tent gave in to the bears.
```

`translate()`에서 `text` 메서드를 호출하고 `translation_table`를 인수로 전달한 다음, 결과를 `encrypted_text`라는 변수에 할당하세요.

# --hints--

`encrypted_text`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

`text.translate(translation_table)`에 `encrypted_text`을 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
--fcc-editable-region--

--fcc-editable-region--
```
