---
id: 680a82806c4c6bc3567f2708
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`text`라는 새 변수를 선언하고 문자열 `hello world`을 할당하세요. 이것이 암호화할 메시지가 됩니다.

# --hints--

`text`라는 변수를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("text")`)) })
```

문자열 `hello world`을 `text`에 할당해야 합니다. 문자열은 작은따옴표나 큰따옴표로 감싸고 대소문자에 주의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("text").is_equivalent("text = 'hello world'")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--

--fcc-editable-region--
```
