---
id: 680a41ef5bfdc059cfd9d4b5
title: Step 2
challengeType: 20
dashedName: step-2
---

# --description--

`alphabet`라는 변수를 하나 더 선언하고, 이 변수에 `abcdefghijklmnopqrstuvwxyz` 문자열을 할당하세요.

# --hints--

`alphabet`라는 변수를 선언해야 합니다. 변수 이름을 줄 맨 앞에 두는 것에 주의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("alphabet")`)) })
```

`abcdefghijklmnopqrstuvwxyz` 변수에 `alphabet` 문자열을 할당해야 합니다. 문자열을 감쌀 때는 작은따옴표나 큰따옴표 중 하나를 사용하세요.

```js
({ test: () => assert(runPython(`
_Node(_code).find_variable("alphabet").is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
shift = 5
```
