---
id: 6810891bd6de8a87a833df42
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

출력 결과에서 볼 수 있듯이 메시지가 암호화되었습니다. 다음 단계는 다른 메시지를 암호화하고 싶을 때 코드를 재사용할 수 있도록 만드는 것입니다.

이를 위해서는 함수를 만들어야 합니다. 참고로, 터미널에 `spam`을 출력하는 `Spam!`라는 함수를 만드는 방법은 다음과 같습니다.

```py
def spam():
    print('Spam!')
```

`caesar`라는 이름의 함수를 만드세요. 기존 코드를 모두 함수 본문 안에 넣으세요. 함수 본문 내 모든 줄의 들여쓰기 수준을 동일하게 유지하는 것에 주의하세요.

# --hints--

`caesar`라는 이름의 함수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("caesar")`)) })
```

지금까지 작성한 모든 코드를 `caesar` 함수 본문 안으로 옮겨야 합니다. 함수 본문 내 모든 줄의 들여쓰기 수준을 동일하게 유지하는지 확인하세요...

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_body().is_equivalent("alphabet = 'abcdefghijklmnopqrstuvwxyz'\\nshift = 5\\nshifted_alphabet = alphabet[shift:] + alphabet[:shift]\\ntranslation_table = str.maketrans(alphabet, shifted_alphabet)\\ntext = 'hello world'\\nencrypted_text = text.translate(translation_table)\\nprint(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
print(encrypted_text)
--fcc-editable-region--
```
