---
id: 6810869355fa2468f9ebaf9c
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이제 터미널에 출력 결과를 확인하려면 `encrypted_text`를 출력하세요.

# --hints--

`encrypted_text`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
alphabet = 'abcdefghijklmnopqrstuvwxyz'
shift = 5
shifted_alphabet = alphabet[shift:] + alphabet[:shift]
translation_table = str.maketrans(alphabet, shifted_alphabet)
text = 'hello world'
encrypted_text = text.translate(translation_table)
--fcc-editable-region--

--fcc-editable-region--
```
