---
id: 6818b80519e6c6c47a8bb6e8
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이전 수업에서 배운 것처럼, 함수는 함수 내에서 참조할 수 있는 변수인 매개변수를 가질 수 있습니다. 다음은 두 개의 매개변수 `sum`과 `num1`를 가진 `num2` 함수입니다.

```py
def sum(num1, num2):
    print(num1 + num2)
```

암호화할 메시지와 시프트 값이 여전히 함수 내에 하드코딩되어 있습니다. 함수에 `text`과 `shift`라는 두 개의 매개변수를 추가하고, `text` 함수 본문에서 `shift`와 `caesar` 변수 선언을 삭제하세요.

# --hints--

`caesar` 함수는 두 개의 매개변수 `text`과 `shift`를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift")`)) })
```

`text` 함수 본문 내에 `caesar` 변수가 선언되어 있으면 안 됩니다.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("text")`)) })
```

`shift` 함수 본문 내에 `caesar` 변수가 선언되어 있으면 안 됩니다.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").find_body().has_variable("shift")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shift = 5
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    text = 'hello world'
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
```
