---
id: 6818eb3f98c0bca5af7c8d03
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

이제 코드를 재사용할 수 있습니다. 하지만 `caesar` 함수는 터미널에 메시지를 출력하고 기본적으로 `None`을 반환합니다. 이를 증명하려면 코드 마지막에 `encrypted_text`를 출력하세요.

# --hints--

함수 본문 밖에서 `encrypted_text`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(encrypted_text)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--

--fcc-editable-region--
```
