---
id: 6818e9c7c8dc5694e8f64fc3
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 `caesar` 함수를 문자열 `freeCodeCamp`과 숫자 `3`를 인수로 호출하여 테스트하세요. 함수 호출 결과를 `encrypted_text`라는 변수에 할당하세요.

# --hints--

전역 범위에 `encrypted_text`라는 변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("encrypted_text")`)) })
```

`caesar('freeCodeCamp', 3)`를 `encrypted_text` 변수에 할당해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = caesar('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
