---
id: 6818ed1adc249bb91d373729
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이전 수업에서 `return` 문이 함수에서 값을 반환하는 데 사용되어 코드의 다른 곳에서 사용할 수 있다는 것을 배웠습니다.

```py
def spam():
    return 'Spam!'
```

함수에서 `print(encrypted_text)` 호출을 삭제하세요. 그런 다음 `encrypted_text` 변수 선언을 삭제하고 함수에서 `text.translate(translation_table)`을 직접 반환하세요.

# --hints--

`print(encrypted_text)` 함수 내에 `caesar`가 없어야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_call("print(encrypted_text)")`)) })
```

`encrypted_text` 함수 내에 `caesar` 변수가 없어야 합니다.

```js
({ test: () => assert(runPython(`not _Node(_code).find_function("caesar").has_variable("encrypted_text")`)) })
```

`caesar` 함수는 `text.translate(translation_table)`을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_return("text.translate(translation_table)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    encrypted_text = text.translate(translation_table)
    print(encrypted_text)
--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
