---
id: 6818fcb250f34e62cab32c39
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

메시지가 암호화되었지만 대문자는 그대로 남아 있는 것을 눈치챘을 수 있습니다. 이는 변환 표에 대문자가 포함되어 있지 않기 때문입니다.

이 문제를 고치려면 `str.maketrans()` 메서드를 사용해 `upper()` 호출에 전달하는 문자열을 수정해야 합니다.

```py
greet = "Hello"
print(greet.upper()) # HELLO
```

각 인수에 그 인수 자체의 대문자 버전을 연결하여 `str.maketrans()` 호출을 업데이트하세요.

# --hints--

각 인수에 그 인수 자체의 대문자 버전을 연결하여 `str.maketrans()` 호출의 인수를 수정해야 합니다.

```js
({ test: () => assert(runPython(`caesar('This Is A TEST', 5) == 'Ymnx Nx F YJXY'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
--fcc-editable-region--
    translation_table = str.maketrans(alphabet, shifted_alphabet)
--fcc-editable-region--
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
