---
id: 6819bc0637b80256a33adccc
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

암호의 기본 기능을 구현했으니 이제 유효성 검사를 추가할 차례입니다. 이를 위해 `if` 문이 필요합니다. `if` 문 구문을 다시 한 번 확인하세요:

```py
if condition:
    # code to run when condition is true
```

함수 본문 시작 부분에 `if` 문을 만드세요. 지금은 조건으로 `True` 을 사용하고, `if` 문 본문 안에서 문자열 `Shift must be an integer value.` 를 반환하세요.

# --hints--

`if` 함수 안에 `caesar` 문이 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0]`)) })
```

`if` 문은 조건으로 `True` 을 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("True")`)) })
```

`Shift must be an integer value.` 문에서 문자열 `if` 을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_bodies()[0].has_return("'Shift must be an integer value.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
