---
id: 6819ca464db219ab86f37dc4
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

이미 확인했듯이, 텍스트를 암호화하는 데 사용되는 시프트는 양의 정수여야 합니다. 하지만 `25`의 마지막 인덱스인 `alphabet`를 초과할 수는 없습니다.

`if` 문에 두 번째 조건을 추가하여 `shift`이 `25`보다 큰지 확인하세요. Python에서 논리 OR 연산은 `or` 연산자를 통해 구현된다는 점을 기억하세요.

또한, 반환되는 메시지를 `Shift must be an integer between 1 and 25.`로 업데이트하세요.

# --hints--

`or` 뒤에 두 번째 조건을 추가할 때 `shift < 1` 연산자를 사용하여 `shift`가 `25`보다 큰지 확인하세요. 기존 조건은 수정하지 마세요.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift > 25", "25 < shift", "shift >= 26", "26 =< shift"]
assert any(_cond.is_equivalent(f"shift < 1 or {option}") for option in _options)
`) })
```

두 번째 `if` 문은 `Shift must be an integer between 1 and 25.`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be an integer between 1 and 25.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    if shift < 1:
        return 'Shift must be a positive integer.'
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
