---
id: 6819c5c9faf2548e556163d5
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`not` 연산자는 불리언 값을 반대로 바꿉니다. `True`를 `False`로, `False`를 `True`로 변경합니다:

```py
print(not True) # False
print(not False) # True
```

`not` 연산자를 사용해 `if` 문장의 조건을 고쳐서, `Shift must be an integer value.`가 정수가 *아닐 때* 함수가 `shift`을 반환하도록 하세요.

# --hints--

`not` 연산자를 사용해 `isinstance(shift, int)`를 `if` 문장에서 부정해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if isinstance(shift, int):
--fcc-editable-region--
        return 'Shift must be an integer value.'
    
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
