---
id: 6819be9055828d65c813eec5
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

`isinstance()` 함수는 첫 번째 인수가 두 번째 인수의 인스턴스이면 `True`을 반환하고, 그렇지 않으면 `False`를 반환합니다:

```py
print(isinstance('Hello World', str)) # True
print(isinstance(42, int)) # True
```

현재 `if` 문의 조건을 `isinstance()` 호출로 바꾸세요. 첫 번째 인수로 `shift`를, 두 번째 인수로 `int`를 전달하세요.

# --hints--

`if` 문은 조건으로 `isinstance(shift, int)`를 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[0].find_conditions()[0].is_equivalent("isinstance(shift, int)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):

--fcc-editable-region--
    if True:
--fcc-editable-region--
        return 'Shift must be an integer value.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
