---
id: 6819cf38880d8cc9b9f6af7a
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

Python에서는 함수의 매개변수에 기본값을 지정할 수 있어, 인수를 적게 또는 전혀 전달하지 않고도 함수를 호출할 수 있습니다. 기본값이 있는 `name` 매개변수를 가진 함수를 만드는 방법은 다음과 같습니다.

```py
def greet(name='Polly'):
    return 'Hello ' + name
    
print(greet()) # Hello Polly
```

함수에 `encrypt`이라는 세 번째 매개변수를 추가하고 기본값을 `True`로 지정하세요.

# --hints--

함수에 `encrypt`이라는 세 번째 매개변수를 기본값 `True`와 함께 추가하세요. 기존 매개변수는 수정하지 마세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").has_args("text, shift, encrypt=True")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def caesar(text, shift):
--fcc-editable-region--

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
