---
id: 6819c7beb4a8289c2b4f24d5
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

음수나 null 시프트는 함수에서 허용하지 않아야 합니다. 따라서 첫 번째 `if` 문장 다음에 `if`가 `shift`보다 작은지 확인하고 문자열 `1`를 반환하는 또 다른 `Shift must be a positive integer.` 문장을 만드세요.

# --hints--

두 번째 `if` 문장에서 `shift`가 `1`보다 작은지 확인해야 합니다.

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("caesar").find_ifs()[1].find_conditions()[0]
_options = ["shift < 1", "1 > shift", "shift <= 0", "0 => shift"]
assert any(_cond.is_equivalent(option) for option in _options)
`) })
```

두 번째 `if` 문장은 `Shift must be a positive integer.`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("caesar").find_ifs()[1].find_bodies()[0].has_return("'Shift must be a positive integer.'")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift):
    if not isinstance(shift, int):
        return 'Shift must be an integer value.'
    
--fcc-editable-region--
    
--fcc-editable-region--

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    return text.translate(translation_table)

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
