---
id: 681a2a6cf95e82ca30866364
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

`encrypt`와 `decrypt`이라는 이름의 두 함수를 선언하세요. 두 함수 모두 `text`와 `shift` 매개변수를 가집니다.

암호화 과정에는 `encrypt`를, 복호화 과정에는 `decrypt`를 사용하며, 두 작업에 설명적인 이름을 붙이세요.

두 새 함수 모두에서 `caesar`과 `text`을 전달하는 `shift` 호출을 반환하되, `False` 함수에는 세 번째 인수로 `decrypt`도 반드시 전달하세요.

# --hints--

`encrypt`라는 이름의 함수가 두 개의 매개변수 `text`와 `shift`를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_args("text, shift")`)) })
```

`encrypt` 함수는 `caesar(text, shift)`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("encrypt").has_return("caesar(text, shift)")`)) })
```

`decrypt`라는 이름의 함수가 두 개의 매개변수 `text`와 `shift`를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_function("decrypt").has_args("text, shift")`)) })
```

`decrypt` 함수는 `caesar(text, shift, False)`를 반환해야 합니다.

```js
({ test: () => runPython(`
_n = _Node(_code).find_function("decrypt")
assert _n.has_return("caesar(text, shift, encrypt=False)") or _n.has_return("caesar(text, shift, False)")
`) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

--fcc-editable-region--

--fcc-editable-region--

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
