---
id: 681a2e35c406ade2f0a5f9b2
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

이제 `encrypt` 함수를 테스트할 시간입니다. 같은 인수를 사용하여 `caesar` 호출을 `encrypt` 호출로 바꾸세요. 터미널에 같은 출력이 나타날 것입니다.

# --hints--

`encrypted_text` 변수는 `encrypt('freeCodeCamp', 3)` 값을 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("encrypted_text").is_equivalent("encrypted_text = encrypt('freeCodeCamp', 3)")`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'

    if not encrypt:
        shift = - shift
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

def encrypt(text, shift):
    return caesar(text, shift)
    
def decrypt(text, shift):
    return caesar(text, shift, encrypt=False)

--fcc-editable-region--
encrypted_text = caesar('freeCodeCamp', 3)
--fcc-editable-region--
print(encrypted_text)
```
