---
id: 681a1d60e037127e7c8e0230
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

이전 단계에서 추가한 매개변수를 사용해 함수가 전달받은 텍스트를 암호화할지(기본 동작, `encrypt=True`) 아니면 암호화된 메시지를 복호화할지 결정하세요.

`if`가 참이 아닌지 확인하는 `encrypt` 문을 만드세요. 새 `if` 문 안에서 `shift`를 `- shift`로 설정하세요. 이는 암호화 과정과 반대 방향으로 시프트가 이루어지도록 하기 위해 필요합니다.

# --hints--

`shift`가 참이 아닐 때 `- shift`를 `encrypt`로 설정해야 합니다.

```js
({ test: () => assert(runPython(`caesar('Ymnx Nx F Yjxy', 5, False) == 'This Is A Test'`)) })
```

# --seed--

## --seed-contents--

```py
def caesar(text, shift, encrypt=True):

    if not isinstance(shift, int):
        return 'Shift must be an integer value.'

    if shift < 1 or shift > 25:
        return 'Shift must be an integer between 1 and 25.'

    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    
--fcc-editable-region--
    
--fcc-editable-region--
    
    shifted_alphabet = alphabet[shift:] + alphabet[:shift]
    translation_table = str.maketrans(alphabet + alphabet.upper(), shifted_alphabet + shifted_alphabet.upper())
    encrypted_text = text.translate(translation_table)
    return encrypted_text

encrypted_text = caesar('freeCodeCamp', 3)
print(encrypted_text)
```
