---
id: 66e45c8140f9fda5c105ae26
title: بناء صفحة منشور حدث
challengeType: 25
dashedName: build-an-event-flyer-page
demoType: onClick
---

# --description--

**الهدف:** أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن يحتوي المستند على عنصر `header` داخل الجسم.
1. يجب أن يحتوي عنصر `header` على صورة لحدثك، ثم عنصر `h1` يحمل عنوان الحدث، بهذا الترتيب. يمكنك استخدام هذه الصورة إذا رغبت: `https://cdn.freecodecamp.org/curriculum/labs/event.jpg`
1. يجب أن يحتوي المستند على عنصر `main` داخل الجسم.
1. يجب أن يحتوي عنصر `main` على عنصرين على الأقل من نوع `section` لعرض ميزات الحدث.
1. يجب أن يحتوي كل عنصر من عناصر `section` على عنصر `h2` بداخله.
1. يجب أن يكون هناك حشو علوي وسفلي للجسم بقيمة `50px`.
1. يجب أن يكون هناك هامش علوي وسفلي للجسم بقيمة `0`، وهامش أيمن وأيسر يحقق التوسيط.
1. يجب أن يكون عرض الجسم مضبوطًا نسبةً إلى عرض إطار العرض.
1. يجب أن يستخدم الجسم دالة `calc` لضبط خاصية `min-height` بحيث تساوي 100% من ارتفاع إطار العرض مطروحًا منه كل الحشو المطبق أعلى وأسفل الجسم.
1. يجب أن يحتوي المنشور على عنصر واحد على الأقل من نوع `hr`.
1. يجب ضبط خاصية `width` لجميع عناصر `hr` و`section` بنسبة مئوية بالنسبة للأصل.

**ملاحظة:** تأكد من ربط ورقة الأنماط في ملف HTML وتطبيق CSS الخاص بك.

# --hints--

يجب أن يكون لديك عنصر `header` داخل عنصر `body`.

```js
assert.exists(document.querySelector('body > header'));
```

يجب أن يكون لديك عنصر `img` كأول عنصر فرعي في `header`.

```js
assert.equal(document.querySelector('header')?.children?.[0].tagName, 'IMG');
```

يجب أن يكون لديك عنصر `h1` كثاني عنصر فرعي في `header`.

```js
assert.equal(document.querySelector('header')?.children?.[1].tagName, 'H1');
```

يجب أن يكون لديك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يحتوي عنصر `h1` على عنوان الحدث الخاص بك.

```js
assert.isAtLeast(document.querySelector('h1').innerText.length, 1);
```

يجب أن يكون لديك عنصر `main` داخل عنصر `body`.

```js
assert.exists(document.querySelector('body > main'));
```

يجب أن يحتوي عنصر `main` على عنصرين `section` على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('main > section').length, 2);
```

يجب أن يحتوي كل عنصر `section` على عنصر `h2` بداخله.

```js
const sections = document.querySelectorAll('main > section');
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(sections.length, 1);
assert.equal(sections.length, headers.length);
```

يجب ألا تكون عناصر `h2` فارغة.

```js
const headers = document.querySelectorAll('main > section > h2');
assert.isAtLeast(headers.length, 1);
headers.forEach(header => assert.notEqual(header.innerHTML.trim(), ''));
```

يجب أن يحتوي عنصر `body` على حشو علوي وسفلي بمقدار `50px`.

```js
const body = window.getComputedStyle(document.querySelector('body'));
assert.equal(body?.paddingTop, '50px');
assert.equal(body?.paddingBottom, '50px');
```

يجب أن يحتوي عنصر `body` على هامش علوي وسفلي بمقدار `0`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginTop, '0px');
assert.equal(body?.marginBottom, '0px');
```

يجب أن يحتوي عنصر `body` على هامش أيسر وأيمن بقيمة `auto`.

```js
const body = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(body?.marginRight, 'auto');
assert.equal(body?.marginLeft, 'auto');
```

يجب أن يضبط عنصر `body` عرضه باستخدام `vw`.

```js
assert.isTrue(new __helpers.CSSHelp(document).getStyle('body')?.width?.endsWith('vw'));
```

يجب أن يستخدم عنصر `body` دالة `calc` لضبط `min-height` إلى `100vh - 100px`.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('body')?.minHeight, ['calc(-100px + 100vh)', 'calc(100vh - 100px)']);
```

يجب أن يكون لديك عنصر `hr` واحد على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('hr')?.length, 1);
```

يجب ضبط عرض عناصر `hr` باستخدام قيمة نسبية مئوية.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'hr' && 
    rule.style.width?.endsWith('%')
  )
));
```

يجب ضبط عرض عناصر `section` باستخدام قيمة نسبية مئوية.

```js
assert.isTrue([...new __helpers.CSSHelp(document).getCSSRules().values()].some(rule => 
  rule.selectorText?.split(',').some(selector => 
    selector.trim() === 'section' && 
    rule.style.width?.endsWith('%')
  )
));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Build an Event Flyer Page</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>

  <header>
    <img src="https://cdn.freecodecamp.org/curriculum/labs/rock.jpg">
    <h1>Pet Rock Festival</h1>
    <p>Join us for the first annual Pet Rock Festival!</p>
    <p><b>When:</b> <time datetime="2025-10-03">October 3rd</time> <b>Where:</b> Boulder Park</p>
  </header>

  <hr>

  <main>
    <section>
      <h2>Festival Highlights</h2>
      <ul>
        <li>Pet Rock Obedience Shows</li>
        <li>Best Dressed Pet Rock Contest</li>
        <li>Pet Rock Uphill Racing</li>
      </ul>
    </section>

    <section>
      <h2>Special Guests</h2>
      <ul>
        <li>John Gravel</li>
        <li>Wendy Cobblestone</li>
      </ul>
    </section>

    <section>
      <h2>Food Trucks</h2>
      <ul>
        <li>Rock 'n' Roll Tacos</li>
        <li>Granite Grill</li>
        <li>Stone Cold Smoothies</li>
      </ul>
    </section>
  </main>

  <hr>

<footer>
  <p>&copy; 2025 Pet Rock Festival</p>
</footer>

</body>

</html>
```

```css
html {
  background-color: gray;
}

body {
  padding: 50px 10px;
  margin: 0 auto;
  width: 90vw;
  min-height: calc(100vh - 100px);
  background-color: white;
  font-family: sans-serif;
  text-align: center;
}

hr {
  margin: 30px auto;
  width: 80%;
  background-color: black;
  height: 2px;
}

main {
  margin: 0 auto;
}

section {
  text-align: center;
  padding: 50px 5px;
  vertical-align: top;
  display: inline-block;
  min-height: 100px;
  width: 30%;
}

ul {
  text-align: left;
  list-style: inside;
  padding-left: 10px;
}
```
