---
id: 69de4308e652aaeb5ccdb03e
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie den Umgang mit den String-Methoden von Python.

Beginnen Sie damit, eine Variable namens `email` zu erstellen und ihr den String `alice.johnson@company.com` zuzuweisen.

# --hints--

Sie sollten eine Variable namens `email` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email')
`) })
```

Sie sollten den String `alice.johnson@company.com` der Variable `email` zuweisen. Denken Sie daran, den Text entweder in einfache oder doppelte Anführungszeichen zu setzen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('email').is_equivalent("email = 'alice.johnson@company.com'")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
