---
id: 69e20a0451c7a3ef189514d3
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Da Sie nun die Position von `@` in der E-Mail kennen, können Sie den Benutzernamen extrahieren, der vor dem `@`-Zeichen steht.

Im vorherigen Workshop haben Sie einen Teil eines Strings mithilfe von Slicing extrahiert. Zur Erinnerung: Die Syntax lautet:

```python
text = 'hello world'
print(text[:5])  # hello
```

Verwenden Sie Slicing auf `email`, um alles vor `at_position` zu extrahieren, und speichern Sie das Ergebnis in einer Variablen namens `username`.

Geben Sie dann `username` aus, sodass die Ausgabe `Username: alice.johnson` lautet.

# --hints--

Sie sollten eine Variable namens `username` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('username')
`) })
```

Sie sollten Slicing mit `at_position` auf `email` verwenden, um `username` einen Wert zuzuweisen.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('username')
assert node.is_equivalent("username = email[:at_position]") or node.is_equivalent("username = email[0:at_position]")
`) })
```

Sie sollten `username` ausgeben, sodass die Ausgabe `Username: alice.johnson` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {username}')",
    "print('Username: ' + username)",
    "print('Username:', username)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

--fcc-editable-region--

--fcc-editable-region--
```
