---
id: 69e20876baed33de636b989a
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Manchmal möchten Sie die Position eines bestimmten Zeichens in einem String finden. Python stellt dafür die eingebaute `.find()`-Methode bereit. Sie gibt eine ganze Zahl zurück, die den Index (die Position) des ersten Vorkommens des angegebenen Teilstrings darstellt.

Hier ist ein Beispiel:

```py
text = 'hello'
print(text.find('e'))  # 1
```

Rufen Sie nun die `.find()`-Methode auf `email` auf und übergeben Sie `'@'`, um die Position des At-Zeichens im `email`-String zu finden. Weisen Sie das Ergebnis einer Variablen namens `at_position` zu.

Geben Sie dann `at_position` aus, sodass die Ausgabe lautet: `Position of @: 13`.

# --hints--

Sie sollten eine Variable namens `at_position` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_position')
`) })
```

Sie sollten `at_position` den Wert `email.find('@')` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_position').is_equivalent("at_position = email.find('@')")
`) })
```

Sie sollten `at_position` ausgeben, sodass die Ausgabe `Position of @: 13` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Position of @: {at_position}')",
    "print('Position of @: ' + str(at_position))",
    "print('Position of @:', at_position)",
    "print('Position of @:', str(at_position))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')
--fcc-editable-region--

--fcc-editable-region--
```
