---
id: 69e6284f3b7659c5a05c2a79
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Sie können überprüfen, ob ein Teilstring in einem String enthalten ist, indem Sie das Schlüsselwort `in` verwenden. Es gibt `True` zurück, wenn der Teilstring vorhanden ist, und `False`, wenn nicht. Hier ist ein Beispiel:

```py
sentence = 'I love Python'
print('Python' in sentence) # True
print('Java' in sentence) # False
```

Erstellen Sie eine Variable namens `has_at`, die überprüft, ob das Zeichen `@` in `email` enthalten ist.

Geben Sie dann `has_at` aus, sodass die Ausgabe lautet: `Has @ symbol: True`.

# --hints--

Sie sollten eine Variable namens `has_at` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_at')
`) })
```

Ihre Variable `has_at` sollte den Wert `'@' in email` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_at').is_equivalent("has_at = '@' in email")
`) })
```

Sie sollten `has_at` ausgeben, sodass die Ausgabe `Has @ symbol: True` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has @ symbol: {has_at}')",
    "print('Has @ symbol: ' + str(has_at))",
    "print('Has @ symbol:', has_at)",
    "print('Has @ symbol:', str(has_at))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

--fcc-editable-region--

--fcc-editable-region--
```
