---
id: 69e62ce96df237884da25282
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Die `.lower()`-Methode wandelt alle Zeichen in einem String in Kleinbuchstaben um. Das ist nützlich, um E-Mail-Adressen zu normalisieren:

```python
upperGreet = 'HELLO WORLD'
lowerGreet = upperGreet.lower()
print(lowerGreet)  # Output: hello world
```

Sie können sowohl `.strip()` als auch `.lower()` auf diese Weise auf einen String aufrufen:

```python
messy = '  HELLO WORLD  '
clean = messy.strip().lower()
print(clean)  # Output: hello world
```

Erstellen Sie eine Variable `messy_email` und weisen Sie ihr den String `SARAH.DAVIS@COMPANY.COM` mit zwei führenden und zwei nachfolgenden Leerzeichen zu.

Rufen Sie außerdem `.strip().lower()` auf `messy_email` auf und weisen Sie das Ergebnis einer Variable `clean_email` zu.

Geben Sie dann beide aus, sodass die Ausgabe lautet:

```md
Original email: '  SARAH.DAVIS@COMPANY.COM  '
Cleaned email: 'sarah.davis@company.com'
```

# --hints--

Sie sollten eine Variable namens `messy_email` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_email')
`) })
```

Sie sollten <code>'&nbsp;&nbsp;SARAH.DAVIS@COMPANY.COM&nbsp;&nbsp;'</code> `messy_email` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_email').is_equivalent("messy_email = '  SARAH.DAVIS@COMPANY.COM  '")
`) })
```

Sie sollten eine Variable namens `clean_email` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_email')
`) })
```

Sie sollten `.strip().lower()` auf `messy_email` aufrufen und das Ergebnis einer Variable `clean_email` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_email').is_equivalent("clean_email = messy_email.strip().lower()")
`) })
```

Sie sollten `messy_email` ausgeben, sodass die Ausgabe `Original email: '  SARAH.DAVIS@COMPANY.COM  '` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original email: '{messy_email}'\\")",
    "print(\\"Original email: '\\" + messy_email + \\"'\\")",
    "print(\\"Original email: '\\", messy_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Sie sollten `clean_email` ausgeben, sodass die Ausgabe `Cleaned email: 'sarah.davis@company.com'` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned email: '{clean_email}'\\")",
    "print(\\"Cleaned email: '\\" + clean_email + \\"'\\")",
    "print(\\"Cleaned email: '\\", clean_email, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
