---
id: 69e6299710319e7cb8ebdc37
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Daten aus der Praxis sind oft unordentlich mit zusätzlichen Leerzeichen und uneinheitlicher Groß- und Kleinschreibung.

Erstellen Sie eine Variable `messy_name` und weisen Sie ihr den String `sARaH dAVis` mit zwei führenden und zwei abschließenden Leerzeichen zu. Geben Sie sie dann so aus, dass die Ausgabe lautet:

```md
Original name: '  sARaH dAVis  '
```

Sie werden das in den nächsten Schritten bereinigen.

# --hints--

Sie sollten eine Variable namens `messy_name` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('messy_name')
`) })
```

Sie sollten <code>'&nbsp;&nbsp;sARaH dAVis&nbsp;&nbsp;'</code> `messy_name` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('messy_name').is_equivalent("messy_name = '  sARaH dAVis  '")
`) })
```

Sie sollten `messy_name` ausgeben, sodass die Ausgabe `Original name: '  sARaH dAVis  '` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Original name: '{messy_name}'\\")",
    "print(\\"Original name: '\\" + messy_name + \\"'\\")",
    "print(\\"Original name: '\\", messy_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

--fcc-editable-region--

--fcc-editable-region--
```
