---
id: 69e736c39dfa403496e8071e
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Die `.upper()`-Methode wandelt alle Zeichen in Großbuchstaben um. Das ist nützlich für Namensschilder oder Überschriften:

```py
name = 'sarah davis'
badge = name.upper()
print(badge)  # Output: SARAH DAVIS
```

Erstellen Sie eine Variable namens `name_badge` und weisen Sie ihr das Ergebnis zu, das durch das Aufrufen der `.upper()`-Methode auf `clean_name` entsteht.

Geben Sie dann `name_badge` aus, sodass die Ausgabe lautet: `Name badge: SARAH DAVIS`.

# --hints--

Sie sollten eine Variable namens `name_badge` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('name_badge')
`) })
```

Sie sollten `.upper()` auf `clean_name` aufrufen und das Ergebnis der Variable `name_badge` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('name_badge').is_equivalent("name_badge = clean_name.upper()")
`) })
```

Sie sollten `name_badge` ausgeben, sodass die Ausgabe `Name badge: SARAH DAVIS` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Name badge: {name_badge}')",
    "print('Name badge: ' + name_badge)",
    "print('Name badge:', name_badge)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

--fcc-editable-region--

--fcc-editable-region--
```
