---
id: 69e73a141addc6f442924b0d
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Sie haben bereits `.find()` und Slicing auf der Variablen `email` verwendet. Jetzt werden Sie dieselbe Technik auf `clean_email` anwenden.

Finden Sie die Position des Zeichens `@` in `clean_email` und speichern Sie sie in `at_pos`. Verwenden Sie dann Slicing, um alles vor dem `@` zu extrahieren, und speichern Sie es in `email_user`.

Geben Sie anschließend `clean_email` und `email_user` aus, sodass die Ausgabe lautet:

```md
Email: sarah.davis@company.com
Username: sarah.davis
```

# --hints--

Sie sollten eine Variable namens `at_pos` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('at_pos')
`) })
```

Sie sollten `.find('@')` auf `clean_email` aufrufen und das Ergebnis der Variablen `at_pos` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('at_pos').is_equivalent("at_pos = clean_email.find('@')")
`) })
```

Sie sollten eine Variable namens `email_user` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('email_user')
`) })
```

Sie sollten Slicing mit `at_pos` auf `clean_email` verwenden und das Ergebnis der Variablen `email_user` zuweisen.

```js
({ test: () => runPython(`
node = _Node(_code).find_variable('email_user')
assert node.is_equivalent("email_user = clean_email[:at_pos]") or node.is_equivalent("email_user = clean_email[0:at_pos]")
`) })
```

Sie sollten `clean_email` ausgeben, sodass die Ausgabe `Email: sarah.davis@company.com` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {clean_email}')",
    "print('Email: ' + clean_email)",
    "print('Email:', clean_email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Sie sollten `email_user` ausgeben, sodass die Ausgabe `Username: sarah.davis` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Username: {email_user}')",
    "print('Username: ' + email_user)",
    "print('Username:', email_user)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

--fcc-editable-region--

--fcc-editable-region--
```
