---
id: 69e75dc335b6ec9903a87d3e
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Sie können auch `.replace()` verwenden, um unerwünschte Zeichen zu entfernen. Zum Beispiel, um Bindestriche aus einer Telefonnummer zu entfernen:

```py
phone = '456-147'
clean = phone.replace('-', '')
print(clean)  # Output: 456147
```

Erstellen Sie eine Variable `phone` und weisen Sie ihr den String `123-456-7890` zu. Ersetzen Sie dann `-` durch einen leeren String und weisen Sie das Ergebnis einer Variablen namens `phone_clean` zu.

Geben Sie dann beide aus, sodass die Ausgabe lautet:

```md
Phone with dashes: 123-456-7890
Phone without dashes: 1234567890
```

# --hints--

Sie sollten eine Variable namens `phone` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone')
`) })
```

Sie sollten den String `123-456-7890` der Variable `phone` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone').is_equivalent("phone = '123-456-7890'")
`) })
```

Sie sollten eine Variable namens `phone_clean` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('phone_clean')
`) })
```

Sie sollten `.replace('-', '')` auf `phone` anwenden und das Ergebnis der Variablen `phone_clean` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('phone_clean').is_equivalent("phone_clean = phone.replace('-', '')")
`) })
```

Sie sollten `phone` ausgeben, sodass die Ausgabe `Phone with dashes: 123-456-7890` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone with dashes: {phone}')",
    "print('Phone with dashes: ' + phone)",
    "print('Phone with dashes:', phone)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

Sie sollten `phone_clean` ausgeben, sodass die Ausgabe `Phone without dashes: 1234567890` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Phone without dashes: {phone_clean}')",
    "print('Phone without dashes: ' + phone_clean)",
    "print('Phone without dashes:', phone_clean)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

--fcc-editable-region--

--fcc-editable-region--
```
