---
id: 69e73f4f5f986657647bbd63
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Die Methode `.replace(old, new)` ersetzt jede Vorkommnis des Substrings `old` durch einen anderen, `new`. Sie können sie auch mit anderen Methoden verketten. Zum Beispiel:

```py
text = 'john.doe'
name = text.replace('.', ' ').title()
print(name)  # Output: John Doe
```

Ersetzen Sie in `email_user` `.` durch ein einzelnes Leerzeichen und rufen Sie dann `.title()` auf dem Ergebnis auf. Weisen Sie das Ergebnis einer Variablen namens `display_name` zu.

Geben Sie dann `display_name` aus, sodass die Ausgabe lautet: `Display name: Sarah Davis`.

# --hints--

Sie sollten eine Variable namens `display_name` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('display_name')
`) })
```

Sie sollten `.replace('.', ' ')` und `.title()` auf `email_user` anwenden und das Ergebnis der Variablen `display_name` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('display_name').is_equivalent("display_name = email_user.replace('.', ' ').title()")
`) })
```

Sie sollten `display_name` ausgeben, sodass die Ausgabe `Display name: Sarah Davis` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Display name: {display_name}')",
    "print('Display name: ' + display_name)",
    "print('Display name:', display_name)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

--fcc-editable-region--

--fcc-editable-region--
```
