---
id: 69e76d1997661be7f0040e2d
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Die `.split(separator)`-Methode teilt einen String in eine Liste von Teilstrings an jeder Stelle, an der `separator` vorkommt. Zum Beispiel:

```py
fruits = 'apple, banana, cherry'
parts = fruits.split(', ')
print(parts)  # Output: ['apple', 'banana', 'cherry']
```

Erstellen Sie eine Variable `full_address` und weisen Sie ihr den String `123 Main Street, Springfield, IL` zu. Erstellen Sie dann `address_parts`, indem Sie an einem Komma gefolgt von einem Leerzeichen aufteilen.

Geben Sie anschließend `address_parts` aus, sodass die Ausgabe lautet: `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

# --hints--

Sie sollten eine Variable namens `full_address` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('full_address')
`) })
```

Sie sollten den String `123 Main Street, Springfield, IL` der Variable `full_address` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('full_address').is_equivalent("full_address = '123 Main Street, Springfield, IL'")
`) })
```

Sie sollten eine Variable namens `address_parts` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('address_parts')
`) })
```

Sie sollten `.split(', ')` auf `full_address` anwenden und das Ergebnis der Variable `address_parts` zuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('address_parts').is_equivalent("address_parts = full_address.split(', ')")
`) })
```

Sie sollten `address_parts` ausgeben, sodass die Ausgabe lautet: `Address parts: ['123 Main Street', 'Springfield', 'IL']`.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Address parts: {address_parts}')",
    "print('Address parts: ' + str(address_parts))",
    "print('Address parts:', address_parts)",
    "print('Address parts:', str(address_parts))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

--fcc-editable-region--

--fcc-editable-region--
```
