---
id: 69eb3158c33d9a229a79951c
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Die Methode `.count(substring)` gibt zurück, wie oft ein Teilstring in einem String vorkommt. Zum Beispiel:

```py
text = 'banana'
print(text.count('a'))  # Output: 3
```

Erstellen Sie zwei Variablen namens `dot_count` und `dash_count`. `dot_count` soll zählen, wie viele `.`-Zeichen in `clean_email` vorkommen, und `dash_count` soll zählen, wie viele `-`-Zeichen in `employee_code` vorkommen.

# --hints--

Sie sollten eine Variable namens `dot_count` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dot_count')
`) })
```

Sie sollten `.count('.')` auf `clean_email` verwenden, um `dot_count` einen Wert zuzuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dot_count').is_equivalent("dot_count = clean_email.count('.')")
`) })
```

Sie sollten eine Variable namens `dash_count` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('dash_count')
`) })
```

Sie sollten `.count('-')` auf `employee_code` verwenden, um `dash_count` einen Wert zuzuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('dash_count').is_equivalent("dash_count = employee_code.count('-')")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

clean_name = messy_name.strip().title()
print(f"Cleaned name: '{clean_name}'")

messy_email = '  SARAH.DAVIS@COMPANY.COM  '
clean_email = messy_email.strip().lower()
print(f"Original email: '{messy_email}'")
print(f"Cleaned email: '{clean_email}'")

name_badge = clean_name.upper()
print(f'Name badge: {name_badge}')

at_pos = clean_email.find('@')
email_user = clean_email[:at_pos]

print(f'Email: {clean_email}')
print(f'Username: {email_user}')

display_name = email_user.replace('.', ' ').title()
print(f'Display name: {display_name}')

phone = '123-456-7890'
phone_clean = phone.replace('-', '')
print(f'Phone with dashes: {phone}')
print(f'Phone without dashes: {phone_clean}')

full_address = '123 Main Street, Springfield, IL'
address_parts = full_address.split(', ')
print(f'Address parts: {address_parts}')

rejoined = ' | '.join(address_parts)
print(f'Rejoined: {rejoined}')

employee_code = 'EMP-2024-SD'
starts_with_emp = employee_code.startswith('EMP')
ends_with_com = clean_email.endswith('.com')

print(f'Code: {employee_code}')
print(f'Starts with EMP: {starts_with_emp}')
print(f'Clean email: {clean_email}')
print(f'Ends with .com: {ends_with_com}')

--fcc-editable-region--

--fcc-editable-region--
```
