---
id: 6a71a8df10e823e77358d896
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Erstellen Sie eine weitere Variable namens `has_com`, die überprüft, ob `.com` in `email` enthalten ist.

Geben Sie dann `has_com` aus, sodass die Ausgabe lautet: `Has .com: True`.

# --hints--

Sie sollten eine Variable namens `has_com` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('has_com')
`) })
```

Ihre Variable `has_com` sollte den Wert `'.com' in email` haben.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('has_com').is_equivalent("has_com = '.com' in email")
`) })
```

Sie sollten `has_com` ausgeben, sodass die Ausgabe `Has .com: True` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Has .com: {has_com}')",
    "print('Has .com: ' + str(has_com))",
    "print('Has .com:', has_com)",
    "print('Has .com:', str(has_com))",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')
--fcc-editable-region--

--fcc-editable-region--
```
