---
id: 6a71b3e49668c3be461257b5
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Die `.title()`-Methode schreibt den ersten Buchstaben jedes Wortes groß. Hier ist ein Beispiel:

```py
lowercase = 'hello world'
title_case = lowercase.title()
print(title_case)  # Output: Hello World
```

Erstellen Sie eine Variable namens `clean_name` und weisen Sie ihr das Ergebnis zu, das durch den Aufruf von `.title()` auf `messy_name` entsteht.

# --hints--

Sie sollten eine Variable namens `clean_name` erstellen.

```js
({ test: () => runPython(`
assert _Node(_code).has_variable('clean_name')
`) })
```

Sie sollten `.title()` auf `messy_name` anwenden, um `clean_name` einen Wert zuzuweisen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.title()")
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--

--fcc-editable-region--
```
