---
id: 6a86b9e64da746838e67fb13
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Die `.strip()`-Methode entfernt Leerzeichen an beiden Enden eines Strings.

```python
messy = '  hello world  '
clean = messy.strip()
print(clean)  # Output: hello world
```

Sie können `.strip()` und `.title()` verketten, um Leerzeichen an beiden Enden eines Strings zu entfernen und den ersten Buchstaben jedes Wortes zu kapitalisieren:

```python
messy = '  hello world  '
clean = messy.strip().title()
print(clean)  # Output: Hello World
```

Aktualisieren Sie den Wert von `clean_name`, indem Sie `.strip()` auf `messy_name` aufrufen, bevor Sie `.title()` aufrufen.

Geben Sie danach `clean_name` aus, sodass die Ausgabe `Cleaned name: 'Sarah Davis'` lautet.

# --hints--

Sie sollten `clean_name` aktualisieren, indem Sie `.strip()` auf `messy_name` aufrufen, bevor Sie `.title()` aufrufen.

```js
({ test: () => runPython(`
assert _Node(_code).find_variable('clean_name').is_equivalent("clean_name = messy_name.strip().title()")
`) })
```

Sie sollten `clean_name` ausgeben, sodass die Ausgabe `Cleaned name: 'Sarah Davis'` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f\\"Cleaned name: '{clean_name}'\\")",
    "print(\\"Cleaned name: '\\" + clean_name + \\"'\\")",
    "print(\\"Cleaned name: '\\", clean_name, \\"'\\", sep='')",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
print(f'Email: {email}')

at_position = email.find('@')
print(f'Position of @: {at_position}')

username = email[:at_position]
print(f'Username: {username}')

has_at = '@' in email
print(f'Has @ symbol: {has_at}')

has_com = '.com' in email
print(f'Has .com: {has_com}')

messy_name = '  sARaH dAVis  '
print(f"Original name: '{messy_name}'")

--fcc-editable-region--
clean_name = messy_name.title()
--fcc-editable-region--
```
