---
id: 6a870cc60c23f0332a2fe788
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Erinnern Sie sich aus einem vorherigen Workshop, dass Sie einen f-String verwenden können, um eine Variable in einen String einzufügen, indem Sie ein `f` vor das öffnende Anführungszeichen schreiben und die Variable in `{}` einschließen.

Hier eine Erinnerung daran, wie es funktioniert:

```python
name = 'Alice'
print(f'Hello, {name}')  # Hello, Alice
```

Geben Sie jetzt `email` aus, sodass die Ausgabe lautet: `Email: alice.johnson@company.com`.

# --hints--

Sie sollten `email` ausgeben, sodass die Ausgabe `Email: alice.johnson@company.com` lautet.

```js
({ test: () => runPython(`
node = _Node(_code)
valid_calls = [
    "print(f'Email: {email}')",
    "print('Email: ' + email)",
    "print('Email:', email)",
]
assert any(node.has_call(call) for call in valid_calls)
`) })
```

# --seed--

## --seed-contents--

```py
email = 'alice.johnson@company.com'
--fcc-editable-region--

--fcc-editable-region--
```
