---
id: 56104e9e514f539506016a5c
title: "`final-expression`で奇数を反復する"
challengeType: 1
forumTopicId: 18212
dashedName: iterate-odd-numbers-with-a-for-loop
---

# --description--

Forループは1ずつ反復する必要はありません。`final-expression`を変えることで、偶数ずつ数えることができます。

`i = 0`から始めて、`i < 10`の間ループします。`i += 2`を使って、`i`をループごとに2ずつ増やします。

```js
const ourArray = [];

for (let i = 0; i < 10; i += 2) {
  ourArray.push(i);
}
```

`ourArray`には今、`[0, 2, 4, 6, 8]`が入っています。`initialization`を変えて、奇数ずつ数えられるようにしましょう。

# --instructions--

`for`ループを使って、1から9までの奇数を`myArray`に追加してください。

# --hints--

ここでは`for`ループを使う必要があります。

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

`myArray`は`[1, 3, 5, 7, 9]`と等しくなるはずです。

```js
assert.deepEqual(myArray, [1, 3, 5, 7, 9]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 1; i < 10; i += 2) {
  myArray.push(i);
}
```
