---
id: 56105e7b514f539506016a5e
title: forループで逆方向にカウントする
challengeType: 1
forumTopicId: 16808
dashedName: count-backwards-with-a-for-loop
---

# --description--

forループは、正しい条件を定義できれば逆方向にもカウントできます。

各繰り返しで2ずつ減らすには、初期化、条件、最後の式を変更する必要があります。

`i = 10`から始めて、`i > 0`の間ループします。`i`で`i -= 2`を2ずつ減らします。

```js
const ourArray = [];

for (let i = 10; i > 0; i -= 2) {
  ourArray.push(i);
}
```

`ourArray`には今や`[10, 8, 6, 4, 2]`が入ります。初期化と最後の式を変えて、2ずつ減る降順の奇数の配列を作りましょう。

# --instructions--

`myArray`ループを使って、9から1までの奇数を`for`にプッシュしてください。

# --hints--

これには`for`ループを使うべきです。

```js
assert(/for\s*\([^)]+?\)/.test(__helpers.removeJSComments(code)));
```

配列のメソッド`push`を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/myArray.push/));
```

`myArray`は`[9, 7, 5, 3, 1]`と等しくなるはずです。

```js
assert.deepEqual(myArray, [9, 7, 5, 3, 1]);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [];

// Only change code below this line

```

# --solutions--

```js
const myArray = [];
for (let i = 9; i > 0; i -= 2) {
  myArray.push(i);
}
```
