---
id: 56533eb9ac21ba0edf2244b8
title: プラスイコール演算子で文字列を連結する
challengeType: 1
forumTopicId: 16803
dashedName: concatenating-strings-with-the-plus-equals-operator
---

# --description--

`+=` 演算子を使って、既存の文字列変数の末尾に文字列を<dfn>連結</dfn>することもできます。これは長い文字列を複数行に分けるのにとても役立ちます。

**注意:** スペースに気をつけてください。連結は連結された文字列の間にスペースを自動で追加しないので、自分でスペースを入れる必要があります。

例:

```js
let ourStr = "I come first. ";
ourStr += "I come second.";
```

`ourStr` は文字列 `I come first. I come second.` の値を持ちます。

# --instructions--

`myStr` を複数行にわたって、`This is the first sentence.` と `This is the second sentence.` の2つの文字列を `+=` 演算子で連結して作成してください。例のように `+=` 演算子を使い、2つの文字列の間にスペースを必ず入れてください。まず最初の文字列を `myStr` に代入し、次に2つ目の文字列を追加してください。

# --hints--

`myStr` には2つの文字列の間に半角スペースが1つ入っている必要があります。

```js
assert(/sentence\. This/.test(myStr));
```

`myStr` は文字列 `This is the first sentence. This is the second sentence.` の値を持つべきです。

```js
assert(myStr === 'This is the first sentence. This is the second sentence.');
```

`+=` 演算子を使って `myStr` を作成してください。

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*(["']).*\1/g));
```

# --seed--

## --seed-contents--

```js
let myStr;
```

# --solutions--

```js
let myStr = "This is the first sentence. ";
myStr += "This is the second sentence.";
```
