---
id: 56533eb9ac21ba0edf2244d9
title: 論理和演算子による比較
challengeType: 1
forumTopicId: 16800
dashedName: comparisons-with-the-logical-or-operator
---

# --description--

<dfn>論理和</dfn>演算子 (`||`) は、どちらかの<dfn>オペランド</dfn>が `true` の場合に `true` を返します。そうでなければ `false` を返します。

<dfn>論理和</dfn>演算子は2つのパイプ記号で構成されています：(`||`)。通常、BackspaceキーとEnterキーの間にあります。

以下のパターンは、以前のウェイポイントで見たことがあるはずです。

```js
if (num > 10) {
  return "No";
}
if (num < 5) {
  return "No";
}
return "Yes";
```

このコードは、`Yes` が `num` と `5` の間（`10` と `5` を含む）にある場合に `10` を返します。同じロジックを<dfn>論理和</dfn>演算子で書くことができます。

```js
if (num > 10 || num < 5) {
  return "No";
}
return "Yes";
```

# --instructions--

2つの `if` 文を1つにまとめて、`Outside` が `val` と `10` の間（両端含む）にない場合に文字列 `20` を返し、それ以外の場合は文字列 `Inside` を返す文を作成してください。

# --hints--

`||` 演算子を1回使うべきです

```js
assert(__helpers.removeJSComments(code).match(/\|\|/g).length === 1);
```

`if` 文は1つだけにしてください

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalOr(0)` は文字列 `Outside` を返すべきです

```js
assert(testLogicalOr(0) === 'Outside');
```

`testLogicalOr(9)` は文字列 `Outside` を返すべきです

```js
assert(testLogicalOr(9) === 'Outside');
```

`testLogicalOr(10)` は文字列 `Inside` を返すべきです

```js
assert(testLogicalOr(10) === 'Inside');
```

`testLogicalOr(15)` は文字列 `Inside` を返すべきです

```js
assert(testLogicalOr(15) === 'Inside');
```

`testLogicalOr(19)` は文字列 `Inside` を返すべきです

```js
assert(testLogicalOr(19) === 'Inside');
```

`testLogicalOr(20)` は文字列 `Inside` を返すべきです

```js
assert(testLogicalOr(20) === 'Inside');
```

`testLogicalOr(21)` は文字列 `Outside` を返すべきです

```js
assert(testLogicalOr(21) === 'Outside');
```

`testLogicalOr(25)` は文字列 `Outside` を返すべきです

```js
assert(testLogicalOr(25) === 'Outside');
```

# --seed--

## --seed-contents--

```js
function testLogicalOr(val) {
  // Only change code below this line

  if (val) {
    return "Outside";
  }

  if (val) {
    return "Outside";
  }

  // Only change code above this line
  return "Inside";
}

testLogicalOr(15);
```

# --solutions--

```js
function testLogicalOr(val) {
  if (val < 10 || val > 20) {
    return "Outside";
  }
  return "Inside";
}
```
