---
id: 56533eb9ac21ba0edf2244b1
title: 複合代入による乗算
challengeType: 1
forumTopicId: 16662
dashedName: compound-assignment-with-augmented-multiplication
---

# --description--

`*=` 演算子は変数に数値を乗算します。

```js
myVar = myVar * 5;
```

は `myVar` に `5` を乗算します。これは次のように書き換えられます。

```js
myVar *= 5;
```

# --instructions--

`a`、`b`、`c` の代入を `*=` 演算子を使うように変換してください。

# --hints--

`a` は `25` と等しくなるはずです。

```js
assert(a === 25);
```

`b` は `36` と等しくなるはずです。

```js
assert(b === 36);
```

`c` は `46` と等しくなるはずです。

```js
assert(c === 46);
```

各変数に `*=` 演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\*=/g).length === 3);
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(
  /let a = 5;/.test(__helpers.removeJSComments(code)) &&
    /let b = 12;/.test(__helpers.removeJSComments(code)) &&
    /let c = 4\.6;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 5;
let b = 12;
let c = 4.6;

// Only change code below this line
a = a * 5;
b = 3 * b;
c = c * 10;
```

# --solutions--

```js
let a = 5;
let b = 12;
let c = 4.6;

a *= 5;
b *= 3;
c *= 10;
```
