---
id: 56533eb9ac21ba0edf2244b7
title: プラス演算子による文字列の連結
challengeType: 1
forumTopicId: 16802
dashedName: concatenating-strings-with-plus-operator
---

# --description--

JavaScriptでは、`+`演算子が`String`値と一緒に使われると、<dfn>連結</dfn>演算子と呼ばれます。複数の文字列を<dfn>連結</dfn>して新しい文字列を作ることができます。

**例**

```js
'My name is Alan,' + ' I concatenate.'
```

**注意:** スペースに気をつけてください。連結は文字列の間にスペースを自動で入れないので、自分でスペースを追加する必要があります。

例:

```js
const ourStr = "I come first. " + "I come second.";
```

文字列`I come first. I come second.`がコンソールに表示されます。
# --instructions--

`myStr`と`This is the start.`の文字列から、`This is the end.`演算子を使って`+`を作ってください。2つの文字列の間にスペースを入れることを忘れないでください。

# --hints--

`myStr`には2つの文字列の間に半角スペースが1つ入っているはずです。

```js
assert(/start\. This/.test(myStr));
```

`myStr`には文字列`This is the start. This is the end.`の値が入っているはずです。

```js
assert(myStr === 'This is the start. This is the end.');
```

`+`を作るには`myStr`演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/(["']).*\1\s*\+\s*(["']).*\2/g));
```

`myStr`は`const`キーワードを使って作成するべきです。

```js
assert(/const\s+myStr/.test(__helpers.removeJSComments(code)));
```

結果は`myStr`変数に代入してください。

```js
assert(/myStr\s*=/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "This is the start. " + "This is the end.";
```
