---
id: 56533eb9ac21ba0edf2244de
title: switch文にデフォルトオプションを追加する
challengeType: 1
forumTopicId: 16653
dashedName: adding-a-default-option-in-switch-statements
---

# --description--

`switch` 文では、すべての可能な値を `case` 文として指定できない場合があります。その代わりに、どの `default` 文にも一致しなかった場合に実行される `case` 文を追加できます。これは `else` チェーンの最後の `if/else` 文のようなものだと考えてください。

`default` 文は最後のケースであるべきです。

```js
switch (num) {
  case value1:
    statement1;
    break;
  case value2:
    statement2;
    break;
...
  default:
    defaultStatement;
    break;
}
```

# --instructions--

次の条件に対して `answer` を設定する switch 文を書いてください：  
`a` - `apple`  
`b` - `bird`  
`c` - `cat`  
`default` - `stuff`

# --hints--

`switchOfStuff("a")` は文字列 `apple` を返すべきです

```js
assert(switchOfStuff('a') === 'apple');
```

`switchOfStuff("b")` は文字列 `bird` を返すべきです

```js
assert(switchOfStuff('b') === 'bird');
```

`switchOfStuff("c")` は文字列 `cat` を返すべきです

```js
assert(switchOfStuff('c') === 'cat');
```

`switchOfStuff("d")` は文字列 `stuff` を返すべきです

```js
assert(switchOfStuff('d') === 'stuff');
```

`switchOfStuff(4)` は文字列 `stuff` を返すべきです

```js
assert(switchOfStuff(4) === 'stuff');
```

`if` や `else` 文は使わないでください

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

`default` 文を使うべきです

```js
assert(switchOfStuff('string-to-trigger-default-case') === 'stuff');
```

少なくとも 3 つの `break` 文が必要です

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function switchOfStuff(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

switchOfStuff(1);
```

# --solutions--

```js
function switchOfStuff(val) {
  let answer = "";

  switch(val) {
    case "a":
      answer = "apple";
      break;
    case "b":
      answer = "bird";
      break;
    case "c":
      answer = "cat";
      break;
    default:
      answer = "stuff";
  }
  return answer;
}
```
