---
id: 56533eb9ac21ba0edf2244b4
title: シングルクォートで文字列を囲む
challengeType: 1
forumTopicId: 18260
dashedName: quoting-strings-with-single-quotes
---

# --description--

<dfn>文字列</dfn>の値は、JavaScriptではシングルクォートか二重引用符のどちらかで書けますが、開始と終了は同じ種類のクォートでなければなりません。ほかのプログラミング言語とは違い、JavaScriptではシングルクォートと二重引用符は同じように動作します。

```js
const doubleQuoteStr = "This is a string"; 
const singleQuoteStr = 'This is also a string';
```

どちらかのクォートを使う理由は、文字列の中で両方を使いたい場合です。例えば、会話文を文字列に保存して、その会話文をクォートで囲みたいときなどです。もう一つの例は、`<a>` タグを属性付きでクォートに入れて、すべてを文字列内に保存したい場合です。

```js
const conversation = 'Finn exclaims to Jake, "Algebraic!"';
```

しかし、外側のクォートと同じ種類のクォートを文字列の中で使う必要があると問題になります。文字列は開始と終了が同じ種類のクォートで囲まれていますが、そのクォートが途中にあると、文字列がそこで終わったとみなされてエラーになります。

```js
const goodStr = 'Jake asks Finn, "Hey, let\'s go on an adventure?"'; 
const badStr = 'Finn responds, "Let's go!"';
```

ここで`badStr`はエラーをスローします。

上の<dfn>goodStr</dfn>では、バックスラッシュ`\`をエスケープ文字として使うことで、両方のクォートを安全に使えます。

**注意:** バックスラッシュ`\`はスラッシュ`/`とは違います。両者は同じ動作をしません。

# --instructions--

与えられた文字列を、開始と終了をシングルクォートにして、エスケープ文字を使わない文字列に変更してください。

今は文字列内の`<a>`タグがすべて二重引用符を使っています。外側のクォートをシングルクォートに変えて、エスケープ文字を取り除く必要があります。

# --hints--

すべてのバックスラッシュ（`\`）を削除してください。

```js
assert(
  !/\\/g.test(__helpers.removeJSComments(code)) &&
    myStr.match(
      '\\s*<a href\\s*=\\s*"http://www.example.com"\\s*target\\s*=\\s*"_blank">\\s*Link\\s*</a>\\s*'
    )
);
```

シングルクォートが2つ`'`、二重引用符が4つ`"`あるはずです。

```js
assert(__helpers.removeJSComments(code).match(/"/g).length === 4 && __helpers.removeJSComments(code).match(/'/g).length === 2);
```

# --seed--

## --seed-contents--

```js
const myStr = "<a href=\"http://www.example.com\" target=\"_blank\">Link</a>";
```

# --solutions--

```js
const myStr = '<a href="http://www.example.com" target="_blank">Link</a>';
```
