---
id: 56533eb9ac21ba0edf2244af
title: 複合代入と加算代入演算子
challengeType: 1
forumTopicId: 16661
dashedName: compound-assignment-with-augmented-addition
---

# --description--

プログラミングでは、変数の内容を変更するために代入を使うことがよくあります。イコール記号の右側はまず評価されることを覚えておいてください。つまり、次のように書けます。

```js
myVar = myVar + 5;
```

これは `5` に `myVar` を加えることを意味します。このパターンは非常に一般的なので、数学的な演算と代入を一度に行う演算子があります。

そのうちの一つが `+=` 演算子です。

```js
let myVar = 1;
myVar += 5;
console.log(myVar);
```

コンソールには `6` が表示されます。

# --instructions--

`a`、`b`、`c` の代入を `+=` 演算子を使う形に変えてください。

# --hints--

`a` は `15` と等しくなるはずです。

```js
assert(a === 15);
```

`b` は `26` と等しくなるはずです。

```js
assert(b === 26);
```

`c` は `19` と等しくなるはずです。

```js
assert(c === 19);
```

各変数には `+=` 演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\+=/g).length === 3);
```

指定されたコメントの上のコードは変更しないでください。

```js
assert(
  /let a = 3;/.test(__helpers.removeJSComments(code)) &&
    /let b = 17;/.test(__helpers.removeJSComments(code)) &&
    /let c = 12;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 3;
let b = 17;
let c = 12;

// Only change code below this line
a = a + 12;
b = 9 + b;
c = c + 7;
```

# --solutions--

```js
let a = 3;
let b = 17;
let c = 12;

a += 12;
b += 9;
c += 7;
```
