---
id: 56533eb9ac21ba0edf2244a9
title: 代入演算子で変数を初期化する
challengeType: 1
forumTopicId: 301171
dashedName: initializing-variables-with-the-assignment-operator
---

# --description--

変数は宣言と同じ行で初期値に<dfn>初期化</dfn>することがよくあります。

```js
var myVar = 0;
```

`myVar` という新しい変数を作成し、初期値として `0` を代入します。

# --instructions--

`a` という変数を `var` で宣言し、値 `9` で初期化してください。

# --hints--

`a` を値 `9` で初期化する必要があります。

```js
assert(/var\s+a\s*=\s*9(\s*;?\s*)$/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
var a = 9;
```
