---
id: 56533eb9ac21ba0edf2244bf
title: ローカルスコープと関数
challengeType: 1
forumTopicId: 18227
dashedName: local-scope-and-functions
---

# --description--

関数内で宣言された変数や関数のパラメータは<dfn>ローカル</dfn>スコープを持ちます。つまり、それらはその関数内でのみ見えます。

ここに `myTest` という関数があり、`loc` というローカル変数があります。

```js
function myTest() {
  const loc = "foo";
  console.log(loc);
}

myTest();
console.log(loc);
```

`myTest()` 関数の呼び出しはコンソールに文字列 `foo` を表示します。`console.log(loc)` 関数の外側にある `myTest` 行はエラーをスローします。なぜなら `loc` は関数の外では定義されていないからです。

# --instructions--

エディターには何が起きているかを確認できる2つの `console.log` があります。コードを書きながらコンソールをチェックして変化を見てください。`myVar` の中にローカル変数 `myLocalScope` を宣言してテストを実行してください。

**注意:** コンソールには `ReferenceError: myVar is not defined` が表示され続けますが、これはテスト失敗の原因にはなりません。

# --hints--

コードにグローバルな `myVar` 変数を含めないでください。

```js
function declared() {
  myVar;
}

assert.throws(declared, ReferenceError);
```

ローカルな `myVar` 変数を追加してください。

```js
assert(
  /functionmyLocalScope\(\)\{.*(var|let|const)myVar[\s\S]*}/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

# --seed--

## --seed-contents--

```js
function myLocalScope() {
  // Only change code below this line

  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```

# --solutions--

```js
function myLocalScope() {
  // Only change code below this line
  let myVar;
  console.log('inside myLocalScope', myVar);
}
myLocalScope();

// Run and check the console
// myVar is not defined outside of myLocalScope
console.log('outside myLocalScope', myVar);
```
