---
id: 56533eb9ac21ba0edf2244aa
title: 初期化されていない変数の理解
challengeType: 1
forumTopicId: 18335
dashedName: understanding-uninitialized-variables
---

# --description--

JavaScriptの変数は宣言されると、初期値が`undefined`になります。`undefined`な変数で数学的な演算を行うと、結果は`NaN`になり、これは<dfn>「Not a Number」</dfn>を意味します。文字列と`undefined`な変数を連結すると、`undefined`の<dfn>文字列</dfn>が得られます。

# --instructions--

3つの変数`a`、`b`、`c`をそれぞれ`5`、`10`、`"I am a"`で初期化して、`undefined`にならないようにしてください。

# --hints--

`a`は定義されていて、最終的な値が`6`である必要があります。

```js
assert(typeof a === 'number' && a === 6);
```

`b`は定義されていて、最終的な値が`15`である必要があります。

```js
assert(typeof b === 'number' && b === 15);
```

`c`は`undefined`を含まず、最終的な値が文字列の`I am a String!`である必要があります。

```js
assert(!/undefined/.test(c) && c === 'I am a String!');
```

指定されたコメントの下のコードは変更しないでください。

```js
assert(
  /a = a \+ 1;/.test(__helpers.removeJSComments(code)) &&
    /b = b \+ 5;/.test(__helpers.removeJSComments(code)) &&
    /c = c \+ " String!";/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
var a;
var b;
var c;
// Only change code above this line

a = a + 1;
b = b + 5;
c = c + " String!";
```

# --solutions--

```js
var a = 5;
var b = 10;
var c = "I am a";
a = a + 1;
b = b + 5;
c = c + " String!";
```
