---
id: 56533eb9ac21ba0edf2244a8
title: 代入演算子で値をストアする
challengeType: 1
forumTopicId: 18310
dashedName: storing-values-with-the-assignment-operator
---

# --description--

JavaScriptでは、<dfn>代入</dfn>演算子 (`=`) を使って変数に値をストアできます。

```js
myVariable = 5;
```

これは`Number`の値`5`を`myVariable`に代入します。

`=`演算子の右側に計算があれば、それらは左側の変数に値が代入される前に実行されます。

```js
var myVar;
myVar = 5;
```

まず、このコードは`myVar`という名前の変数を作成します。次に、コードは`5`を`myVar`に代入します。これで、コード内で`myVar`が再び現れた場合、プログラムはそれを`5`として扱います。

# --instructions--

変数`7`に値`a`を代入してください。

# --hints--

指定されたコメントより上のコードは変更しないでください。

```js
assert(/var a;/.test(__helpers.removeJSComments(code)));
```

`a`は7の値を持つ必要があります。

```js
assert(typeof a === 'number' && a === 7);
```

# --seed--

## --seed-contents--

```js
// Setup
var a;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
```
