---
id: 56533eb9ac21ba0edf2244dc
title: If Else文のチェーン
challengeType: 1
forumTopicId: 16772
dashedName: chaining-if-else-statements
---

# --description--

`if/else`文は複雑なロジックのためにチェーンでつなげることができます。以下は複数のチェーンされた`if` / `else if`文の<dfn>疑似コード</dfn>です：

```js
if (condition1) {
  statement1
} else if (condition2) {
  statement2
} else if (condition3) {
  statement3
. . .
} else {
  statementN
}
```

# --instructions--

次の条件を満たすチェーンされた`if`/`else if`文を書いてください：

`num < 5` - `Tiny`を返す  
`num < 10` - `Small`を返す  
`num < 15` - `Medium`を返す  
`num < 20` - `Large`を返す  
`num >= 20` - `Huge`を返す

# --hints--

少なくとも4つの`else`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 3);
```

少なくとも4つの`if`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 3);
```

少なくとも1つの`return`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/return/g).length >= 1);
```

`testSize(0)`は文字列`Tiny`を返すべきです

```js
assert(testSize(0) === 'Tiny');
```

`testSize(4)`は文字列`Tiny`を返すべきです

```js
assert(testSize(4) === 'Tiny');
```

`testSize(5)`は文字列`Small`を返すべきです

```js
assert(testSize(5) === 'Small');
```

`testSize(8)`は文字列`Small`を返すべきです

```js
assert(testSize(8) === 'Small');
```

`testSize(10)`は文字列`Medium`を返すべきです

```js
assert(testSize(10) === 'Medium');
```

`testSize(14)`は文字列`Medium`を返すべきです

```js
assert(testSize(14) === 'Medium');
```

`testSize(15)`は文字列`Large`を返すべきです

```js
assert(testSize(15) === 'Large');
```

`testSize(17)`は文字列`Large`を返すべきです

```js
assert(testSize(17) === 'Large');
```

`testSize(20)`は文字列`Huge`を返すべきです

```js
assert(testSize(20) === 'Huge');
```

`testSize(25)`は文字列`Huge`を返すべきです

```js
assert(testSize(25) === 'Huge');
```

# --seed--

## --seed-contents--

```js
function testSize(num) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

testSize(7);
```

# --solutions--

```js
function testSize(num) {
  if (num < 5) {
    return "Tiny";
  } else if (num < 10) {
    return "Small";
  } else if (num < 15) {
    return "Medium";
  } else if (num < 20) {
    return "Large";
  } else {
    return "Huge";
  }
}
```
