---
id: 56533eb9ac21ba0edf2244c0
title: 関数におけるグローバルスコープとローカルスコープの違い
challengeType: 1
forumTopicId: 18194
dashedName: global-vs--local-scope-in-functions
---

# --description--

同じ名前の<dfn>ローカル</dfn>変数と<dfn>グローバル</dfn>変数を両方持つことができます。この場合、ローカル変数がグローバル変数より優先されます。

この例では：

```js
const someVar = "Hat";

function myFun() {
  const someVar = "Head";
  return someVar;
}
```

関数`myFun`は、ローカル変数が存在するため、文字列`Head`を返します。

# --instructions--

`myOutfit`関数にローカル変数を追加して、`outerWear`の値を文字列`sweater`でオーバーライドしてください。

# --hints--

グローバル`outerWear`の値は変更しないでください。

```js
assert(outerWear === 'T-Shirt');
```

`myOutfit`は文字列`sweater`を返すべきです。

```js
assert(myOutfit() === 'sweater');
```

return文は変更しないでください。

```js
assert(/return outerWear/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const outerWear = "T-Shirt";

function myOutfit() {
  // Only change code below this line

  // Only change code above this line
  return outerWear;
}

myOutfit();
```

# --solutions--

```js
const outerWear = "T-Shirt";
function myOutfit() {
  const outerWear = "sweater";
  return outerWear;
}
```
