---
id: 56533eb9ac21ba0edf2244df
title: switch文における複数の同一オプション
challengeType: 1
forumTopicId: 18242
dashedName: multiple-identical-options-in-switch-statements
---

# --description--

もし`break`文が`switch`文の`case`から省略されると、`case`に出会うまで次の`break`文が実行されます。同じ出力を持つ複数の入力がある場合は、次のように`switch`文で表すことができます。

```js
let result = "";
switch (val) {
  case 1:
  case 2:
  case 3:
    result = "1, 2, or 3";
    break;
  case 4:
    result = "4 alone";
}
```

1、2、3のケースはすべて同じ結果を出します。

# --instructions--

次の範囲に対して`answer`を設定するswitch文を書いてください。  
`1-3` - `Low`  
`4-6` - `Mid`  
`7-9` - `High`

**注意:** 範囲内の各数字に対して`case`文が必要です。

# --hints--

`sequentialSizes(1)`は文字列`Low`を返すべきです

```js
assert(sequentialSizes(1) === 'Low');
```

`sequentialSizes(2)`は文字列`Low`を返すべきです

```js
assert(sequentialSizes(2) === 'Low');
```

`sequentialSizes(3)`は文字列`Low`を返すべきです

```js
assert(sequentialSizes(3) === 'Low');
```

`sequentialSizes(4)`は文字列`Mid`を返すべきです

```js
assert(sequentialSizes(4) === 'Mid');
```

`sequentialSizes(5)`は文字列`Mid`を返すべきです

```js
assert(sequentialSizes(5) === 'Mid');
```

`sequentialSizes(6)`は文字列`Mid`を返すべきです

```js
assert(sequentialSizes(6) === 'Mid');
```

`sequentialSizes(7)`は文字列`High`を返すべきです

```js
assert(sequentialSizes(7) === 'High');
```

`sequentialSizes(8)`は文字列`High`を返すべきです

```js
assert(sequentialSizes(8) === 'High');
```

`sequentialSizes(9)`は文字列`High`を返すべきです

```js
assert(sequentialSizes(9) === 'High');
```

`if`や`else`文は使わないでください

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

`case`文は9つ必要です

```js
assert(__helpers.removeJSComments(code).match(/case/g).length === 9);
```

# --seed--

## --seed-contents--

```js
function sequentialSizes(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

sequentialSizes(1);
```

# --solutions--

```js
function sequentialSizes(val) {
  let answer = "";

  switch (val) {
    case 1:
    case 2:
    case 3:
      answer = "Low";
      break;
    case 4:
    case 5:
    case 6:
      answer = "Mid";
      break;
    case 7:
    case 8:
    case 9:
      answer = "High";
  }

  return answer;
}
```
