---
id: 56533eb9ac21ba0edf2244da
title: Else文の紹介
challengeType: 1
forumTopicId: 18207
dashedName: introducing-else-statements
---

# --description--

`if`文の条件が真の場合、その後のコードブロックが実行されます。では、その条件が偽の場合はどうなるでしょうか？通常は何も起こりません。`else`文を使うと、代わりのコードブロックを実行できます。

```js
if (num > 10) {
  return "Bigger than 10";
} else {
  return "10 or Less";
}
```

# --instructions--

`if`文を1つの`if/else`文にまとめてください。

# --hints--

エディター内には`if`文は1つだけにしてください

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`else`文を使うべきです

```js
assert(/else/g.test(__helpers.removeJSComments(code)));
```

`testElse(4)`は文字列`5 or Smaller`を返すべきです

```js
assert(testElse(4) === '5 or Smaller');
```

`testElse(5)`は文字列`5 or Smaller`を返すべきです

```js
assert(testElse(5) === '5 or Smaller');
```

`testElse(6)`は文字列`Bigger than 5`を返すべきです

```js
assert(testElse(6) === 'Bigger than 5');
```

`testElse(10)`は文字列`Bigger than 5`を返すべきです

```js
assert(testElse(10) === 'Bigger than 5');
```

指定されたコメントの上または下のコードは変更しないでください。

```js
assert(/let result = "";/.test(__helpers.removeJSComments(code)) && /return result;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function testElse(val) {
  let result = "";
  // Only change code below this line

  if (val > 5) {
    result = "Bigger than 5";
  }

  if (val <= 5) {
    result = "5 or Smaller";
  }

  // Only change code above this line
  return result;
}

testElse(4);
```

# --solutions--

```js
function testElse(val) {
  let result = "";
  if(val > 5) {
    result = "Bigger than 5";
  } else {
    result = "5 or Smaller";
  }
  return result;
}
```
