---
id: 56533eb9ac21ba0edf2244cd
title: ネストされた配列へのアクセス
challengeType: 1
forumTopicId: 16160
dashedName: accessing-nested-arrays
---

# --description--

これまでの例で見てきたように、オブジェクトはネストされたオブジェクトとネストされた配列の両方を含むことができます。ネストされたオブジェクトにアクセスするのと同様に、配列のブラケット記法を連結してネストされた配列にアクセスできます。

ネストされた配列にアクセスする方法の例を示します。

```js
const ourPets = [
  {
    animalType: "cat",
    names: [
      "Meowzer",
      "Fluffy",
      "Kit-Cat"
    ]
  },
  {
    animalType: "dog",
    names: [
      "Spot",
      "Bowser",
      "Frankie"
    ]
  }
];

ourPets[0].names[1];
ourPets[1].names[0];
```

`ourPets[0].names[1]` は文字列 `Fluffy` であり、`ourPets[1].names[0]` は文字列 `Spot` になります。

# --instructions--

ドット記法とブラケット記法を使って、変数 `secondTree` に `list` 配列の2番目のオブジェクトの `myPlants` 配列の2番目の要素を設定してください。

# --hints--

`secondTree` は文字列 `pine` と等しくなるはずです。

```js
assert(secondTree === 'pine');
```

コードはドット記法とブラケット記法を使って `myPlants` にアクセスする必要があります。

```js
assert(/=\s*myPlants\[1\].list\[1\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = "";
```

# --solutions--

```js
const myPlants = [
  {
    type: "flowers",
    list: [
      "rose",
      "tulip",
      "dandelion"
    ]
  },
  {
    type: "trees",
    list: [
      "fir",
      "pine",
      "birch"
    ]
  }
];

const secondTree = myPlants[1].list[1];
```
