---
id: 56533eb9ac21ba0edf2244ed
title: 文字列に変数を追加する
challengeType: 1
forumTopicId: 16656
dashedName: appending-variables-to-strings
---

# --description--

文字列の<dfn>リテラル</dfn>を複数行にわたって組み立てられるのと同様に、プラスイコール (`+=`) 演算子を使って変数を文字列に追加することもできます。

例:

```js
const anAdjective = "awesome!";
let ourStr = "freeCodeCamp is ";
ourStr += anAdjective;
```

`ourStr` は `freeCodeCamp is awesome!` の値を持ちます。

# --instructions--

`someAdjective` に3文字以上の文字列を設定し、`myStr` に `+=` 演算子を使って追加してください。

# --hints--

`someAdjective` は3文字以上の文字列に設定してください。

```js
assert(typeof someAdjective !== 'undefined' && someAdjective.length > 2);
```

`someAdjective` を `myStr` に `+=` 演算子で追加する必要があります。

```js
assert(__helpers.removeJSComments(code).match(/myStr\s*\+=\s*someAdjective\s*/).length > 0);
```

# --seed--

## --seed-contents--

```js
// Change code below this line
const someAdjective = "";
let myStr = "Learning to code is ";
```

# --solutions--

```js
const someAdjective = "neat";
let myStr = "Learning to code is ";
myStr += someAdjective;
```
