---
id: 56533eb9ac21ba0edf2244ad
title: JavaScriptで数値を1減らす
challengeType: 1
forumTopicId: 17558
dashedName: decrement-a-number-with-javascript
---

# --description--

`--`演算子を使うと、変数の値を簡単に1減らす（デクリメント）ことができます。

```js
i--;
```

は次と同じ意味です。

```js
i = i - 1;
```

**注意：** この行全体が`i--;`となり、等号は不要です。

# --instructions--

コードを変更して、`--`に`myVar`演算子を使ってください。

# --hints--

`myVar`は`10`と等しくなるはずです。

```js
assert(myVar === 10);
```

`myVar = myVar - 1;`を変更してください。

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*myVar\s*[-]\s*1.*?;?/));
```

`myVar`に`10`を代入してはいけません。

```js
assert(!__helpers.removeJSComments(code).match(/myVar\s*=\s*10.*?;?/));
```

`--`に`myVar`演算子を使うべきです。

```js
assert(/[-]{2}\s*myVar|myVar\s*[-]{2}/.test(__helpers.removeJSComments(code)));
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(/let myVar = 11;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 11;

// Only change code below this line
myVar = myVar - 1;
```

# --solutions--

```js
let myVar = 11;
myVar--;
```
