---
id: 56533eb9ac21ba0edf2244bc
title: ショッピングリスト
challengeType: 1
forumTopicId: 18280
dashedName: shopping-list
---

# --description--

変数`myList`にショッピングリストを作成してください。このリストは複数のサブ配列を含む多次元配列にしてください。

各サブ配列の最初の要素は、アイテム名を表す文字列である必要があります。2番目の要素は数量を表す数字である必要があります。例えば、

```js
["Chocolate Bar", 15]
```

リストには少なくとも5つのサブ配列が含まれている必要があります。

# --hints--

`myList`は配列である必要があります。

```js
assert(Array.isArray(myList));
```

各サブ配列の最初の要素はすべて文字列である必要があります。

```js
assert(myList.every(elem => elem && typeof elem[0] === 'string'));
```

各サブ配列の2番目の要素はすべて数字である必要があります。

```js
assert(myList.every(elem => elem && typeof elem[1] === 'number'));
```

リストには少なくとも5つのアイテムが含まれている必要があります。

```js
assert(myList.length > 4);
```

# --seed--

## --seed-contents--

```js
const myList = [];
```

# --solutions--

```js
const myList = [
  ["Candy", 10],
  ["Potatoes", 12],
  ["Eggs", 12],
  ["Catfood", 1],
  ["Toads", 9]
];
```
