---
id: 56533eb9ac21ba0edf2244c3
title: 戻り値を使った代入
challengeType: 1
forumTopicId: 16658
dashedName: assignment-with-a-returned-value
---

# --description--

<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/storing-values-with-the-assignment-operator" target="_blank" rel="noopener noreferrer nofollow">代入演算子を使った値の保存</a>の説明を思い出してください。等号の右側は、値が代入される前にすべて評価されます。つまり、関数の戻り値を変数に代入することができます。

2つの数値を足し合わせる`sum`という関数が定義されているとします。

```js
ourSum = sum(5, 12);
```

`sum`と`5`という引数で`12`関数を呼び出すと、戻り値は`17`になります。この戻り値が`ourSum`変数に代入されます。

# --instructions--

`processArg`関数を`7`という引数で呼び出し、その戻り値を`processed`変数に代入してください。

# --hints--

`processed`は`2`の値を持つべきです

```js
assert(processed === 2);
```

`processArg`を`processed`に代入してください

```js
assert(/processed\s*=\s*processArg\(\s*7\s*\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

// Only change code below this line

```

# --solutions--

```js
var processed = 0;

function processArg(num) {
  return (num + 3) / 5;
}

processed = processArg(7);
```
