---
id: 56533eb9ac21ba0edf2244ca
title: ルックアップにオブジェクトを使う
challengeType: 1
forumTopicId: 18373
dashedName: using-objects-for-lookups
---

# --description--

オブジェクトは辞書のようなキーと値のストレージとして考えられます。表形式のデータがある場合、`switch`文や`if/else`チェーンの代わりにオブジェクトを使って値をルックアップできます。これは入力データが特定の範囲に限定されているとわかっている場合に特に便利です。

以下は article オブジェクトの例です。

```js
const article = {
  "title": "How to create objects in JavaScript",
  "link": "https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/",
  "author": "Kaashan Hussain",
  "language": "JavaScript",
  "tags": "TECHNOLOGY",
  "createdAt": "NOVEMBER 28, 2018"
};

const articleAuthor = article["author"];
const articleLink = article["link"];

const value = "title";
const valueLookup = article[value];
```

`articleAuthor` は文字列 `Kaashan Hussain`、`articleLink` は文字列 `https://www.freecodecamp.org/news/a-complete-guide-to-creating-objects-in-javascript-b0e2450655e8/`、`valueLookup` は文字列 `How to create objects in JavaScript` です。

# --instructions--

switch文を `lookup` というオブジェクトに変換してください。それを使って `val` をルックアップし、対応する文字列を `result` 変数に代入してください。

# --hints--

`phoneticLookup("alpha")` は文字列 `Adams` と等しくなるはずです

```js
assert(phoneticLookup('alpha') === 'Adams');
```

`phoneticLookup("bravo")` は文字列 `Boston` と等しくなるはずです

```js
assert(phoneticLookup('bravo') === 'Boston');
```

`phoneticLookup("charlie")` は文字列 `Chicago` と等しくなるはずです

```js
assert(phoneticLookup('charlie') === 'Chicago');
```

`phoneticLookup("delta")` は文字列 `Denver` と等しくなるはずです

```js
assert(phoneticLookup('delta') === 'Denver');
```

`phoneticLookup("echo")` は文字列 `Easy` と等しくなるはずです

```js
assert(phoneticLookup('echo') === 'Easy');
```

`phoneticLookup("foxtrot")` は文字列 `Frank` と等しくなるはずです

```js
assert(phoneticLookup('foxtrot') === 'Frank');
```

`phoneticLookup("")` は `undefined` と等しくなるはずです

```js
assert(typeof phoneticLookup('') === 'undefined');
```

`return` 文は変更しないでください

```js
assert(__helpers.removeJSComments(code).match(/return\sresult;/));
```

`case`、`switch`、または `if` 文は使わないでください

```js
assert(
  !/case|switch|if/g.test(__helpers.removeJSComments(code).replace(/([/]{2}.*)|([/][*][^/*]*[*][/])/g, ''))
);
```

# --seed--

## --seed-contents--

```js
// Setup
function phoneticLookup(val) {
  let result = "";

  // Only change code below this line
  switch(val) {
    case "alpha":
      result = "Adams";
      break;
    case "bravo":
      result = "Boston";
      break;
    case "charlie":
      result = "Chicago";
      break;
    case "delta":
      result = "Denver";
      break;
    case "echo":
      result = "Easy";
      break;
    case "foxtrot":
      result = "Frank";
  }

  // Only change code above this line
  return result;
}

phoneticLookup("charlie");
```

# --solutions--

```js
function phoneticLookup(val) {
  let result = "";

  const lookup = {
    alpha: "Adams",
    bravo: "Boston",
    charlie: "Chicago",
    delta: "Denver",
    echo: "Easy",
    foxtrot: "Frank"
  };

  result = lookup[val];

  return result;
}
```
