---
id: 56533eb9ac21ba0edf2244c6
title: 並ぶ
challengeType: 1
forumTopicId: 18307
dashedName: stand-in-line
---

# --description--

コンピューターサイエンスにおいて、<dfn>queue</dfn>はアイテムが順序通りに保持される抽象的な<dfn>データ構造</dfn>です。新しいアイテムはキューの後ろに追加され、古いアイテムはキューの前から取り出されます。

# --instructions--

配列（`nextInLine`）と数値（`arr`）を引数に取る関数`item`を作成してください。

数値を配列の末尾に追加し、その後配列の最初の要素を削除してください。

`nextInLine`関数は削除された要素を返す必要があります。

# --hints--

`nextInLine([], 5)`は数値を返す必要があります。

```js
assert.isNumber(nextInLine([], 5));
```

`nextInLine([], 1)`は`1`を返す必要があります。

```js
assert(nextInLine([], 1) === 1);
```

`nextInLine([2], 1)`は`2`を返す必要があります。

```js
assert(nextInLine([2], 1) === 2);
```

`nextInLine([5,6,7,8,9], 1)`は`5`を返す必要があります。

```js
assert(nextInLine([5, 6, 7, 8, 9], 1) === 5);
```

# --seed--

## --seed-contents--

```js
function nextInLine(arr, item) {
  // Only change code below this line
  
  return item;
  // Only change code above this line
}

// Setup
let testArr = [1, 2, 3, 4, 5];

// Display code
console.log("Before: " + JSON.stringify(testArr));
console.log(nextInLine(testArr, 6));
console.log("After: " + JSON.stringify(testArr));
```

# --solutions--

```js
let testArr = [1, 2, 3, 4, 5];

function nextInLine(arr, item) {
    arr.push(item);
    return arr.shift();
}
```
