---
id: 56533eb9ac21ba0edf2244bb
title: ワードブランク
challengeType: 1
forumTopicId: 18377
dashedName: word-blanks
---

# --description--

名詞、動詞、形容詞、副詞などのいくつかの単語が抜けている文が与えられます。抜けている部分に自分の好きな単語を入れて、完成した文が意味をなすようにしてください。

この文を考えてみましょう：

```md
It was really ____, and we ____ ourselves ____.
```

この文には形容詞、動詞、副詞の3つの抜けている部分があり、好きな単語を入れて完成させることができます。完成した文は次のように変数に代入できます：

```js
const sentence = "It was really " + "hot" + ", and we " + "laughed" + " ourselves " + "silly" + ".";
```

# --instructions--

このチャレンジでは、名詞、動詞、形容詞、副詞が与えられます。与えられた単語と自分の好きな単語を使って、完全な文を作ってください。

文字列の連結演算子 `+` を使って、新しい文字列を作成する必要があります。使用する変数は `myNoun`、`myAdjective`、`myVerb`、`myAdverb` です。作成した文字列は `wordBlanks` 変数に代入してください。変数に代入された単語は変更しないでください。

また、文の単語の間にスペースが入るように文字列を作成してください。結果は完全な文になるはずです。

# --before-each--

```js
const removeAssignments = str => str
  .replace(/myNoun\s*=\s*["']dog["']/g, '')
  .replace(/myAdjective\s*=\s*["']big["']/g, '')
  .replace(/myVerb\s*=\s*["']ran["']/g, '')
  .replace(/myAdverb\s*=\s*["']quickly["']/g, '');
```

# --hints--

`wordBlanks` は文字列である必要があります。

```js
assert(typeof wordBlanks === 'string');
```

`myNoun`、`myVerb`、`myAdjective`、`myAdverb` に代入された値は変更しないでください。

```js
assert(
  myNoun === 'dog' &&
    myVerb === 'ran' &&
    myAdjective === 'big' &&
    myAdverb === 'quickly'
);
```

`dog`、`ran`、`big`、`quickly` の値を直接使って `wordBlanks` を作成しないでください。

```js
const newCode = removeAssignments(__helpers.removeJSComments(code));
assert(
  !/dog/.test(newCode) &&
    !/ran/.test(newCode) &&
    !/big/.test(newCode) &&
    !/quickly/.test(newCode)
);
```

`wordBlanks` には、`myNoun`、`myVerb`、`myAdjective`、`myAdverb` に代入されたすべての単語が、単語以外の文字で区切られて（および自分の好きな追加の単語も含めて）含まれている必要があります。

```js
assert(
  /\bdog\b/.test(wordBlanks) &&
    /\bbig\b/.test(wordBlanks) &&
    /\bran\b/.test(wordBlanks) &&
    /\bquickly\b/.test(wordBlanks)
);
```

# --seed--

## --seed-contents--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

// Only change code below this line
const wordBlanks = ""; // Change this line
// Only change code above this line
```

# --solutions--

```js
const myNoun = "dog";
const myAdjective = "big";
const myVerb = "ran";
const myAdverb = "quickly";

let wordBlanks = "Once there was a " + myNoun + " which was very " + myAdjective + ". ";
wordBlanks += "It " + myVerb + " " + myAdverb + " around the yard.";
```
