---
id: 56533eb9ac21ba0edf2244c9
title: 変数を使ったオブジェクトのプロパティへのアクセス
challengeType: 1
forumTopicId: 16165
dashedName: accessing-object-properties-with-variables
---

# --description--

オブジェクトに対するブラケット記法のもう一つの使い方は、変数の値として格納されているプロパティにアクセスすることです。これはオブジェクトのプロパティを繰り返し処理したり、ルックアップテーブルにアクセスしたりするときに非常に便利です。

変数を使ってプロパティにアクセスする例を示します。

```js
const dogs = {
  Fido: "Mutt",
  Hunter: "Doberman",
  Snoopie: "Beagle"
};

const myDog = "Hunter";
const myBreed = dogs[myDog];
console.log(myBreed);
```

文字列 `Doberman` がコンソールに表示されます。

変数名の周りに引用符を使わないことに注意してください。これは、プロパティにアクセスするときに変数の*名前*ではなく*値*を使っているためです。

# --instructions--

`playerNumber` 変数に `16` を代入してください。次に、その変数を使ってプレイヤーの名前を取得し、`player` に代入してください。

# --hints--

`playerNumber` は数字であるべきです

```js
assert(typeof playerNumber === 'number');
```

変数 `player` は文字列であるべきです

```js
assert(typeof player === 'string');
```

`player` の値は文字列 `Montana` であるべきです

```js
assert(player === 'Montana');
```

ブラケット記法を使って `testObj` にアクセスしてください

```js
assert(/testObj\s*?\[.*?\]/.test(__helpers.removeJSComments(code)));
```

変数 `Montana` に値 `player` を直接代入してはいけません

```js
assert(!__helpers.removeJSComments(code).match(/player\s*=\s*"|\'\s*Montana\s*"|\'\s*;/gi));
```

ブラケット記法の中で変数 `playerNumber` を使うべきです

```js
assert(/testObj\s*?\[\s*playerNumber\s*\]/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};

// Only change code below this line
const playerNumber = 42;  // Change this line
const player = testObj;   // Change this line
```

# --solutions--

```js
const testObj = {
  12: "Namath",
  16: "Montana",
  19: "Unitas"
};
const playerNumber = 16;
const player = testObj[playerNumber];
```
