---
id: 56533eb9ac21ba0edf2244b0
title: 複合代入による減算の拡張
challengeType: 1
forumTopicId: 16660
dashedName: compound-assignment-with-augmented-subtraction
---

# --description--

`+=` 演算子と同様に、`-=` は変数から数値を減算します。

```js
myVar = myVar - 5;
```

は `5` から `myVar` を減算します。これは次のように書き換えられます。

```js
myVar -= 5;
```

# --instructions--

`a`、`b`、`c` の代入を `-=` 演算子を使うように変換してください。

# --hints--

`a` は `5` と等しくなるはずです。

```js
assert(a === 5);
```

`b` は `-6` と等しくなるはずです。

```js
assert(b === -6);
```

`c` は `2` と等しくなるはずです。

```js
assert(c === 2);
```

各変数に対して `-=` 演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/-=/g).length === 3);
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(
  /let a = 11;/.test(__helpers.removeJSComments(code)) && /let b = 9;/.test(__helpers.removeJSComments(code)) && /let c = 3;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 11;
let b = 9;
let c = 3;

// Only change code below this line
a = a - 6;
b = b - 15;
c = c - 1;
```

# --solutions--

```js
let a = 11;
let b = 9;
let c = 3;

a -= 6;
b -= 15;
c -= 1;
```
