---
id: 56533eb9ac21ba0edf2244ba
title: 文字列の不変性を理解する
challengeType: 1
forumTopicId: 18331
dashedName: understand-string-immutability
---

# --description--

JavaScriptでは、`String`の値は<dfn>不変</dfn>であり、一度作成されると変更できません。

例えば、次のコードはエラーになります。なぜなら、文字列`B`の中の文字`Bob`を文字`J`に変更できないからです。

```js
let myStr = "Bob";
myStr[0] = "J";
```

ただし、これは`myStr`に再代入できないという意味ではありません。`myStr`を変更する唯一の方法は、新しい値を代入することです。例えば次のようにします。

```js
let myStr = "Bob";
myStr = "Job";
```

# --instructions--

`myStr`への代入を修正して、上の例の方法を使い`Hello World`の文字列値を含むようにしてください。

# --hints--

`myStr`は文字列`Hello World`の値を持つべきです。

```js
assert(myStr === 'Hello World');
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(/myStr = "Jello World"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
let myStr = "Jello World";

// Only change code below this line
myStr[0] = "H"; // Change this line
// Only change code above this line
```

# --solutions--

```js
let myStr = "Jello World";
myStr = "Hello World";
```
