---
id: 56533eb9ac21ba0edf2244d7
title: 以下演算子による比較
challengeType: 1
forumTopicId: 16788
dashedName: comparison-with-the-less-than-or-equal-to-operator
---

# --description--

以下演算子 (`<=`) は2つの数値の値を比較します。左側の数値が右側の数値以下であれば、`true` を返します。左側の数値が右側の数値より大きい場合は、`false` を返します。等値演算子と同様に、以下演算子はデータ型を変換します。

**例**

```js
4   <= 5 // true
'7' <= 7 // true
5   <= 5 // true
3   <= 2 // false
'8' <= 4 // false
```

# --instructions--

return文が意味をなすように、指定された行に以下演算子を追加してください。

# --hints--

`testLessOrEqual(0)` は文字列 `Smaller Than or Equal to 12` を返すべきです

```js
assert(testLessOrEqual(0) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(11)` は文字列 `Smaller Than or Equal to 12` を返すべきです

```js
assert(testLessOrEqual(11) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(12)` は文字列 `Smaller Than or Equal to 12` を返すべきです

```js
assert(testLessOrEqual(12) === 'Smaller Than or Equal to 12');
```

`testLessOrEqual(23)` は文字列 `Smaller Than or Equal to 24` を返すべきです

```js
assert(testLessOrEqual(23) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(24)` は文字列 `Smaller Than or Equal to 24` を返すべきです

```js
assert(testLessOrEqual(24) === 'Smaller Than or Equal to 24');
```

`testLessOrEqual(25)` は文字列 `More Than 24` を返すべきです

```js
assert(testLessOrEqual(25) === 'More Than 24');
```

`testLessOrEqual(55)` は文字列 `More Than 24` を返すべきです

```js
assert(testLessOrEqual(55) === 'More Than 24');
```

以下演算子 `<=` を少なくとも2回使うべきです

```js
assert(__helpers.removeJSComments(code).match(/val\s*<=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessOrEqual(val) {
  if (val) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}

testLessOrEqual(10);
```

# --solutions--

```js
function testLessOrEqual(val) {
  if (val <= 12) {  // Change this line
    return "Smaller Than or Equal to 12";
  }

  if (val <= 24) {  // Change this line
    return "Smaller Than or Equal to 24";
  }

  return "More Than 24";
}
```
