---
id: 56533eb9ac21ba0edf2244ae
title: JavaScriptで余りを求める
challengeType: 1
forumTopicId: 18184
dashedName: finding-a-remainder-in-javascript
---

# --description--

<dfn>余り</dfn>演算子 `%` は、2つの数の割り算の余りを返します。

**例**

<pre>
5 % 2 = 1
5 / 2 = 2 余り 1
2 * 2 = 4
5 - 4 = 1
</pre>

**使い方**  
数学では、数を`2`で割った余りを調べることで、その数が偶数か奇数かを判定できます。偶数は余りが`0`で、奇数は余りが`1`です。

<pre>
17 % 2 = 1
48 % 2 = 0
</pre>

**注意:** <dfn>余り</dfn>演算子は、時に誤ってモジュラス演算子と呼ばれることがあります。非常に似ていますが、負の数に対しては正しく動作しません。

# --instructions--

`remainder`を、`11`を`3`で割った余りに<dfn>余り</dfn>（`%`）演算子を使って設定してください。

# --hints--

変数`remainder`は初期化されている必要があります

```js
assert(/(const|let|var)\s+?remainder/.test(__helpers.removeJSComments(code)));
```

`remainder`の値は`2`であるべきです

```js
assert(remainder === 2);
```

`%`演算子を使うべきです

```js
assert(/\s+?remainder\s*?=\s*?.*%.*;?/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const remainder = 0;
```

# --solutions--

```js
const remainder = 11 % 3;
```
