---
id: 56533eb9ac21ba0edf2244b2
title: 複合代入：拡張除算演算子
challengeType: 1
forumTopicId: 16659
dashedName: compound-assignment-with-augmented-division
---

# --description--

`/=` 演算子は変数を別の数で割ります。

```js
myVar = myVar / 5;
```

`myVar` を `5` で割ります。これは次のように書き換えられます：

```js
myVar /= 5;
```

# --instructions--

`a`、`b`、`c` の代入を `/=` 演算子を使う形に変えてください。

# --hints--

`a` は `4` と等しくなるはずです。

```js
assert(a === 4);
```

`b` は `27` と等しくなるはずです。

```js
assert(b === 27);
```

`c` は `3` と等しくなるはずです。

```js
assert(c === 3);
```

各変数に対して `/=` 演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/\/=/g).length === 3);
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(
  /let a = 48;/.test(__helpers.removeJSComments(code)) &&
    /let b = 108;/.test(__helpers.removeJSComments(code)) &&
    /let c = 33;/.test(__helpers.removeJSComments(code))
);
```

# --seed--

## --seed-contents--

```js
let a = 48;
let b = 108;
let c = 33;

// Only change code below this line
a = a / 12;
b = b / 4;
c = c / 11;
```

# --solutions--

```js
let a = 48;
let b = 108;
let c = 33;

a /= 12;
b /= 4;
c /= 11;
```
