---
id: 56533eb9ac21ba0edf2244d4
title: 大なり演算子による比較
challengeType: 1
forumTopicId: 16786
dashedName: comparison-with-the-greater-than-operator
---

# --description--

大なり演算子 (`>`) は2つの数値の値を比較します。左側の数が右側の数より大きい場合、`true` を返します。そうでなければ、`false` を返します。

等値演算子と同様に、大なり演算子も比較時に値のデータ型を変換します。

**例**

```js
5   >  3  // true
7   > '3' // true
2   >  3  // false
'1' >  9  // false
```

# --instructions--

return文が正しくなるように、指定された行に大なり演算子を追加してください。

# --hints--

`testGreaterThan(0)` は文字列 `10 or Under` を返すべきです

```js
assert(testGreaterThan(0) === '10 or Under');
```

`testGreaterThan(10)` は文字列 `10 or Under` を返すべきです

```js
assert(testGreaterThan(10) === '10 or Under');
```

`testGreaterThan(11)` は文字列 `Over 10` を返すべきです

```js
assert(testGreaterThan(11) === 'Over 10');
```

`testGreaterThan(99)` は文字列 `Over 10` を返すべきです

```js
assert(testGreaterThan(99) === 'Over 10');
```

`testGreaterThan(100)` は文字列 `Over 10` を返すべきです

```js
assert(testGreaterThan(100) === 'Over 10');
```

`testGreaterThan(101)` は文字列 `Over 100` を返すべきです

```js
assert(testGreaterThan(101) === 'Over 100');
```

`testGreaterThan(150)` は文字列 `Over 100` を返すべきです

```js
assert(testGreaterThan(150) === 'Over 100');
```

`>` 演算子を少なくとも2回使う必要があります

```js
assert(__helpers.removeJSComments(code).match(/val\s*>\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterThan(val) {
  if (val) {  // Change this line
    return "Over 100";
  }

  if (val) {  // Change this line
    return "Over 10";
  }

  return "10 or Under";
}

testGreaterThan(10);
```

# --solutions--

```js
function testGreaterThan(val) {
  if (val > 100) {  // Change this line
    return "Over 100";
  }
  if (val > 10) {  // Change this line
    return "Over 10";
  }
  return "10 or Under";
}
```
