---
id: 56533eb9ac21ba0edf2244c8
title: ブラケット記法によるオブジェクトのプロパティへのアクセス
challengeType: 1
forumTopicId: 16163
dashedName: accessing-object-properties-with-bracket-notation
---

# --description--

オブジェクトのプロパティにアクセスする2つ目の方法はブラケット記法(`[]`)です。アクセスしようとしているオブジェクトのプロパティ名にスペースが含まれている場合は、ブラケット記法を使う必要があります。

ただし、スペースがないオブジェクトのプロパティにもブラケット記法は使えます。

以下はブラケット記法を使ってオブジェクトのプロパティを読み取る例です。

```js
const myObj = {
  "Space Name": "Kirk",
  "More Space": "Spock",
  "NoSpace": "USS Enterprise"
};

myObj["Space Name"];
myObj['More Space'];
myObj["NoSpace"];
```

`myObj["Space Name"]`は文字列`Kirk`、`myObj['More Space']`は文字列`Spock`、`myObj["NoSpace"]`は文字列`USS Enterprise`になります。

プロパティ名にスペースが含まれている場合は、シングルクォートかダブルクォートで囲む必要があることに注意してください。

# --instructions--

ブラケット記法を使って`an entree`のプロパティ`the drink`と`testObj`の値を読み取り、それぞれ`entreeValue`と`drinkValue`に代入してください。

# --hints--

`entreeValue`は文字列である必要があります

```js
assert(typeof entreeValue === 'string');
```

`entreeValue`の値は文字列`hamburger`である必要があります

```js
assert(entreeValue === 'hamburger');
```

`drinkValue`は文字列である必要があります

```js
assert(typeof drinkValue === 'string');
```

`drinkValue`の値は文字列`water`である必要があります

```js
assert(drinkValue === 'water');
```

ブラケット記法を2回使う必要があります

```js
assert(__helpers.removeJSComments(code).match(/testObj\s*?\[('|")[^'"]+\1\]/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};

// Only change code below this line
const entreeValue = testObj;   // Change this line
const drinkValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  "an entree": "hamburger",
  "my side": "veggies",
  "the drink": "water"
};
const entreeValue = testObj["an entree"];
const drinkValue = testObj['the drink'];
```
