---
id: 56533eb9ac21ba0edf2244c7
title: ドット記法によるオブジェクトのプロパティへのアクセス
challengeType: 1
forumTopicId: 16164
dashedName: accessing-object-properties-with-dot-notation
---

# --description--

オブジェクトのプロパティにアクセスする方法は2つあります：ドット記法（`.`）と配列に似たブラケット記法（`[]`）です。

ドット記法は、アクセスしたいプロパティの名前が事前にわかっている場合に使います。

以下は、ドット記法（`.`）を使ってオブジェクトのプロパティを読み取る例です：

```js
const myObj = {
  prop1: "val1",
  prop2: "val2"
};

const prop1val = myObj.prop1;
const prop2val = myObj.prop2;
```

`prop1val` は文字列 `val1` の値を持ち、`prop2val` は文字列 `val2` の値を持ちます。

# --instructions--

ドット記法を使って `testObj` のプロパティ値を読み取ってください。変数 `hatValue` にオブジェクトのプロパティ `hat` の値を代入し、変数 `shirtValue` にオブジェクトのプロパティ `shirt` の値を代入してください。

# --hints--

`hatValue` は文字列であるべきです

```js
assert(typeof hatValue === 'string');
```

`hatValue` の値は文字列 `ballcap` であるべきです

```js
assert(hatValue === 'ballcap');
```

`shirtValue` は文字列であるべきです

```js
assert(typeof shirtValue === 'string');
```

`shirtValue` の値は文字列 `jersey` であるべきです

```js
assert(shirtValue === 'jersey');
```

ドット記法を2回使うべきです

```js
assert(__helpers.removeJSComments(code).match(/testObj\.\w+/g).length > 1);
```

# --seed--

## --seed-contents--

```js
// Setup
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

// Only change code below this line
const hatValue = testObj;      // Change this line
const shirtValue = testObj;    // Change this line
```

# --solutions--

```js
const testObj = {
  hat: "ballcap",
  shirt: "jersey",
  shoes: "cleats"
};

const hatValue = testObj.hat;
const shirtValue = testObj.shirt;
```
