---
id: 56533eb9ac21ba0edf2244d1
title: 厳密等値演算子による比較
challengeType: 1
forumTopicId: 16790
dashedName: comparison-with-the-strict-equality-operator
---

# --description--

厳密等値演算子 (`===`) は等値演算子 (`==`) の対応物です。ただし、比較される両方の値を共通の型に変換しようとする等値演算子とは異なり、厳密等値演算子は型変換を行いません。

比較される値の型が異なる場合、それらは等しくないとみなされ、厳密等値演算子は false を返します。

**例**

```js
3 ===  3  // true
3 === '3' // false
```

2つ目の例では、`3` は `Number` 型で、`'3'` は `String` 型です。

# --instructions--

関数が `if` と `Equal` が厳密に等しい場合に文字列 `val` を返すように、`7` 文で厳密等値演算子を使ってください。

# --hints--

`testStrict(10)` は文字列 `Not Equal` を返すべきです

```js
assert(testStrict(10) === 'Not Equal');
```

`testStrict(7)` は文字列 `Equal` を返すべきです

```js
assert(testStrict(7) === 'Equal');
```

`testStrict("7")` は文字列 `Not Equal` を返すべきです

```js
assert(testStrict('7') === 'Not Equal');
```

`===` 演算子を使うべきです

```js
assert(__helpers.removeJSComments(code).match(/(val\s*===\s*\d+)|(\d+\s*===\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrict(val) {
  if (val) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

testStrict(10);
```

# --solutions--

```js
function testStrict(val) {
  if (val === 7) {
    return "Equal";
  }
  return "Not Equal";
}
```
