---
id: 56533eb9ac21ba0edf2244b9
title: 変数を使った文字列の作成
challengeType: 1
forumTopicId: 16805
dashedName: constructing-strings-with-variables
---

# --description--

文字列を作成する必要がある場合があります。連結演算子 (`+`) を使うと、作成中の文字列に1つ以上の変数を挿入できます。

例：

```js
const ourName = "freeCodeCamp";
const ourStr = "Hello, our name is " + ourName + ", how are you?";
```

`ourStr` は文字列 `Hello, our name is freeCodeCamp, how are you?` の値を持ちます。

# --instructions--

`myName` を自分の名前と等しい文字列に設定し、`myStr` と `myName` の間に `My name is ` を挟んで ` and I am well!` を作成してください。

# --hints--

`myName` は少なくとも3文字以上の文字列に設定してください。

```js
assert(typeof myName !== 'undefined' && myName.length > 2);
```

`+` を含む `myStr` を作成するには、`myName` 演算子を2つ使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/["']\s*\+\s*myName\s*\+\s*["']/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const myName = "";
const myStr = "";
```

# --solutions--

```js
const myName = "Bob";
const myStr = "My name is " + myName + " and I am well!";
```
