---
id: 56533eb9ac21ba0edf2244b6
title: 文字列のエスケープシーケンス
challengeType: 1
forumTopicId: 17567
dashedName: escape-sequences-in-strings
---

# --description--

文字列の中でエスケープできるのは引用符だけではありません。エスケープシーケンスを使うと、文字列内で通常は使えない文字を使うことができます。

<table><thead><tr><th>コード</th><th>出力</th></tr></thead><tbody><tr><td><code>\'</code></td><td>シングルクォート</td></tr><tr><td><code>\"</code></td><td>ダブルクォート</td></tr><tr><td><code>\\</code></td><td>バックスラッシュ</td></tr><tr><td><code>\n</code></td><td>改行</td></tr><tr><td><code>\t</code></td><td>タブ</td></tr><tr><td><code>\r</code></td><td>キャリッジリターン</td></tr><tr><td><code>\b</code></td><td>バックスペース</td></tr><tr><td><code>\f</code></td><td>フォームフィード</td></tr></tbody></table>

*バックスラッシュ自体を表示するには、バックスラッシュをエスケープする必要があることに注意してください。*

# --instructions--

次の3行のテキストを、エスケープシーケンスを使って1つの変数`myStr`に代入してください。

<pre>
FirstLine
    \SecondLine
ThirdLine
</pre>

特殊文字を正しく挿入するにはエスケープシーケンスを使う必要があります。また、上記のようなスペースの配置を守り、エスケープシーケンスや単語の間にスペースを入れないでください。

**注意:** `SecondLine`のインデントはスペースではなく、タブのエスケープ文字で実現してください。

# --hints--

`myStr`にはスペースを含めないでください

```js
assert(!/ /.test(myStr));
```

`myStr`には`FirstLine`、`SecondLine`、`ThirdLine`の文字列が含まれている必要があります（大文字小文字を区別してください）

```js
assert(
  /FirstLine/.test(myStr) && /SecondLine/.test(myStr) && /ThirdLine/.test(myStr)
);
```

`FirstLine`の後には改行文字`\n`が続く必要があります

```js
assert(/FirstLine\n/.test(myStr));
```

`myStr`には改行文字の後にタブ文字`\t`が含まれている必要があります

```js
assert(/\n\t/.test(myStr));
```

`SecondLine`の前にはバックスラッシュ文字`\`が必要です

```js
assert(/\\SecondLine/.test(myStr));
```

`SecondLine`と`ThirdLine`の間には改行文字が必要です

```js
assert(/SecondLine\nThirdLine/.test(myStr));
```

`myStr`には指示に示された文字だけを含めてください

```js
assert(myStr === 'FirstLine\n\t\\SecondLine\nThirdLine');
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "FirstLine\n\t\\SecondLine\nThirdLine";
```
