---
id: 56533eb9ac21ba0edf2244db
title: Else If文の紹介
challengeType: 1
forumTopicId: 18206
dashedName: introducing-else-if-statements
---

# --description--

複数の条件を処理する必要がある場合、`if`文と`else if`文を連結して使うことができます。

```js
if (num > 15) {
  return "Bigger than 15";
} else if (num < 5) {
  return "Smaller than 5";
} else {
  return "Between 5 and 15";
}
```

# --instructions--

ロジックを`else if`文を使うように変換してください。

# --hints--

少なくとも2つの`else`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/else/g).length > 1);
```

少なくとも2つの`if`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/if/g).length > 1);
```

各`if else`コードブロックには開きと閉じの波括弧が必要です。

```js
assert(
  __helpers.removeJSComments(code).match(
    /if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s+if\s*\((.+)\)\s*\{[\s\S]+\}\s*else\s*\{[\s\S]+\s*\}/
  )
);
```

`testElseIf(0)`は文字列`Smaller than 5`を返す必要があります

```js
assert(testElseIf(0) === 'Smaller than 5');
```

`testElseIf(5)`は文字列`Between 5 and 10`を返す必要があります

```js
assert(testElseIf(5) === 'Between 5 and 10');
```

`testElseIf(7)`は文字列`Between 5 and 10`を返す必要があります

```js
assert(testElseIf(7) === 'Between 5 and 10');
```

`testElseIf(10)`は文字列`Between 5 and 10`を返す必要があります

```js
assert(testElseIf(10) === 'Between 5 and 10');
```

`testElseIf(12)`は文字列`Greater than 10`を返す必要があります

```js
assert(testElseIf(12) === 'Greater than 10');
```

# --seed--

## --seed-contents--

```js
function testElseIf(val) {
  if (val > 10) {
    return "Greater than 10";
  }

  if (val < 5) {
    return "Smaller than 5";
  }

  return "Between 5 and 10";
}

testElseIf(7);
```

# --solutions--

```js
function testElseIf(val) {
  if(val > 10) {
    return "Greater than 10";
  } else if(val < 5) {
    return "Smaller than 5";
  } else {
    return "Between 5 and 10";
  }
}
```
