---
id: 56533eb9ac21ba0edf2244ac
title: JavaScriptで数値をインクリメントする
challengeType: 1
forumTopicId: 18201
dashedName: increment-a-number-with-javascript
---

# --description--

`++` 演算子を使うと、変数に1を足す、つまり<dfn>インクリメント</dfn>するのが簡単にできます。

```js
i++;
```

は次と同じ意味です。

```js
i = i + 1;
```

**注意：** この行全体が `i++;` となり、等号は不要になります。

# --instructions--

`++` に `myVar` 演算子を使うようにコードを変更してください。

# --hints--

`myVar` は `88` と等しくなるはずです。

```js
assert(myVar === 88);
```

代入演算子は使わないでください。

```js
assert(
  /let\s+myVar\s*=\s*87;\s*\/*.*\s*([+]{2}\s*myVar|myVar\s*[+]{2})/.test(__helpers.removeJSComments(code))
);
```

`++` 演算子を使うべきです。

```js
assert(/[+]{2}\s*myVar|myVar\s*[+]{2}/.test(__helpers.removeJSComments(code)));
```

指定されたコメントより上のコードは変更しないでください。

```js
assert(/let myVar = 87;/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let myVar = 87;

// Only change code below this line
myVar = myVar + 1;
```

# --solutions--

```js
let myVar = 87;
myVar++;
```
