---
id: 56533eb9ac21ba0edf2244c2
title: Return a Value from a Function with Return
challengeType: 1
forumTopicId: 18271
dashedName: return-a-value-from-a-function-with-return
---

# --description--

関数には<dfn>引数</dfn>を渡すことができます。`return`文を使って、関数から値を返すことができます。

**例**

```js
function plusThree(num) {
  return num + 3;
}

const answer = plusThree(5);
```

`answer`は`8`の値を持っています。

`plusThree`は`num`の<dfn>引数</dfn>を受け取り、`num + 3`と等しい値を返します。

# --instructions--

1つの引数を受け取り、それに`timesFive`を掛けて新しい値を返す関数`5`を作成してください。

# --hints--

`timesFive`は関数であるべきです

```js
assert(typeof timesFive === 'function');
```

`timesFive(5)`は`25`を返すべきです

```js
assert(timesFive(5) === 25);
```

`timesFive(2)`は`10`を返すべきです

```js
assert(timesFive(2) === 10);
```

`timesFive(0)`は`0`を返すべきです

```js
assert(timesFive(0) === 0);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function timesFive(num) {
  return num * 5;
}
timesFive(10);
```
