---
id: 56533eb9ac21ba0edf2244e1
title: ネストした For ループ
challengeType: 1
forumTopicId: 18248
dashedName: nesting-for-loops
---

# --description--

多次元配列がある場合、前のウェイポイントと同じロジックを使って、配列とそのサブ配列の両方をループ処理できます。例を示します：

```js
const arr = [
  [1, 2], [3, 4], [5, 6]
];

for (let i = 0; i < arr.length; i++) {
  for (let j = 0; j < arr[i].length; j++) {
    console.log(arr[i][j]);
  }
}
```

これは `arr` の各サブ要素を一度に1つずつ出力します。内側のループでは、`.length` 自体が配列なので、`arr[i]` の `arr[i]` をチェックしていることに注意してください。

# --instructions--

関数 `multiplyAll` を修正して、`arr` のサブ配列内のすべての数字の積を返すようにしてください。

# --hints--

`multiplyAll([[1], [2], [3]])` は `6` を返すべきです

```js
assert(multiplyAll([[1], [2], [3]]) === 6);
```

`multiplyAll([[1, 2], [3, 4], [5, 6, 7]])` は `5040` を返すべきです

```js
assert(
  multiplyAll([
    [1, 2],
    [3, 4],
    [5, 6, 7]
  ]) === 5040
);
```

`multiplyAll([[5, 1], [0.2, 4, 0.5], [3, 9]])` は `54` を返すべきです

```js
assert(
  multiplyAll([
    [5, 1],
    [0.2, 4, 0.5],
    [3, 9]
  ]) === 54
);
```

# --seed--

## --seed-contents--

```js
function multiplyAll(arr) {
  let product = 1;
  // Only change code below this line

  // Only change code above this line
  return product;
}

multiplyAll([[1, 2], [3, 4], [5, 6, 7]]);
```

# --solutions--

```js
function multiplyAll(arr) {
  let product = 1;
  for (let i = 0; i < arr.length; i++) {
    for (let j = 0; j < arr[i].length; j++) {
      product *= arr[i][j];
    }
  }
  return product;
}
```
