---
id: 56533eb9ac21ba0edf2244d5
title: 「Greater Than Or Equal To」演算子による比較
challengeType: 1
forumTopicId: 16785
dashedName: comparison-with-the-greater-than-or-equal-to-operator
---

# --description--

Greater Than Or Equal To 演算子（`>=`）は、2つの数値の値を比較します。左側の数値が右側の数値以上であれば、`true`を返します。そうでなければ、`false`を返します。

等値演算子と同様に、Greater Than Or Equal To 演算子は比較時にデータ型を変換します。

**例**

```js
6   >=  6  // true
7   >= '3' // true
2   >=  3  // false
'7' >=  9  // false
```

# --instructions--

指定された行にGreater Than Or Equal To 演算子を追加して、return文が意味を成すようにしてください。

# --hints--

`testGreaterOrEqual(0)`は文字列`Less than 10`を返すべきです

```js
assert(testGreaterOrEqual(0) === 'Less than 10');
```

`testGreaterOrEqual(9)`は文字列`Less than 10`を返すべきです

```js
assert(testGreaterOrEqual(9) === 'Less than 10');
```

`testGreaterOrEqual(10)`は文字列`10 or Over`を返すべきです

```js
assert(testGreaterOrEqual(10) === '10 or Over');
```

`testGreaterOrEqual(11)`は文字列`10 or Over`を返すべきです

```js
assert(testGreaterOrEqual(11) === '10 or Over');
```

`testGreaterOrEqual(19)`は文字列`10 or Over`を返すべきです

```js
assert(testGreaterOrEqual(19) === '10 or Over');
```

`testGreaterOrEqual(100)`は文字列`20 or Over`を返すべきです

```js
assert(testGreaterOrEqual(100) === '20 or Over');
```

`testGreaterOrEqual(21)`は文字列`20 or Over`を返すべきです

```js
assert(testGreaterOrEqual(21) === '20 or Over');
```

`>=`演算子を少なくとも2回使う必要があります

```js
assert(__helpers.removeJSComments(code).match(/val\s*>=\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testGreaterOrEqual(val) {
  if (val) {  // Change this line
    return "20 or Over";
  }

  if (val) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}

testGreaterOrEqual(10);
```

# --solutions--

```js
function testGreaterOrEqual(val) {
  if (val >= 20) {  // Change this line
    return "20 or Over";
  }

  if (val >= 10) {  // Change this line
    return "10 or Over";
  }

  return "Less than 10";
}
```
