---
id: 56533eb9ac21ba0edf2244cb
title: 複雑なオブジェクトの操作
challengeType: 1
forumTopicId: 18208
dashedName: manipulating-complex-objects
---

# --description--

時には柔軟な<dfn>データ構造</dfn>にデータを格納したい場合があります。JavaScriptのオブジェクトは柔軟なデータを扱う方法の一つです。オブジェクトは任意の組み合わせの<dfn>文字列</dfn>、<dfn>数値</dfn>、<dfn>ブール値</dfn>、<dfn>配列</dfn>、<dfn>関数</dfn>、および<dfn>オブジェクト</dfn>を許容します。

以下は複雑なデータ構造の例です：

```js
const ourMusic = [
  {
    "artist": "Daft Punk",
    "title": "Homework",
    "release_year": 1997,
    "formats": [ 
      "CD", 
      "Cassette", 
      "LP"
    ],
    "gold": true
  }
];
```

これは1つのオブジェクトを含む配列です。そのオブジェクトはアルバムに関するさまざまな<dfn>メタデータ</dfn>を持っています。また、ネストされた`formats`配列も含まれています。もしアルバムのレコードを追加したい場合は、トップレベルの配列にレコードを追加すればよいです。オブジェクトはキーと値の形式のプロパティにデータを保持します。上の例では、`"artist": "Daft Punk"`は`artist`というキーと`Daft Punk`という値を持つプロパティです。

**注意:** 配列内の各オブジェクトの後にはカンマを付ける必要があります。ただし、配列の最後のオブジェクトにはカンマを付けません。

# --instructions--

`myMusic`配列に新しいアルバムを追加してください。`artist`と`title`の文字列、`release_year`の数値、そして文字列の`formats`配列を追加してください。

# --hints--

`myMusic`は配列であるべきです

```js
assert(Array.isArray(myMusic));
```

`myMusic`は少なくとも2つの要素を持つべきです

```js
assert(myMusic.length > 1);
```

`myMusic`配列の要素はオブジェクトであるべきです

```js
myMusic.forEach(object => {assert.typeOf(object, 'object')})
```

`myMusic`のオブジェクトは少なくとも4つのプロパティを持つべきです

```js
myMusic.forEach(object => {assert(Object.keys(object).length > 3); });
```

`myMusic`のオブジェクトは文字列の`artist`プロパティを含むべきです

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['artist']);
  assert.typeOf(object.artist, 'string')
})
```

`myMusic`のオブジェクトは文字列の`title`プロパティを含むべきです

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['title']);
  assert.typeOf(object.title, 'string')
})
```

`myMusic`のオブジェクトは数値の`release_year`プロパティを含むべきです

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['release_year']);
  assert.typeOf(object.release_year, 'number')
})
```

`myMusic`のオブジェクトは配列の`formats`プロパティを含むべきです

```js
myMusic.forEach(object => {
  assert.containsAllKeys(object, ['formats']);
  assert.typeOf(object.formats, 'array')
})
```

`formats`は少なくとも2つの要素を持つ文字列の配列であるべきです

```js
myMusic.forEach(object => {
  object.formats.forEach(format => {
    assert.typeOf(format, 'string')
  });
  assert.isAtLeast(object.formats.length, 2)
})
```

# --seed--

## --seed-contents--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CD",
      "8T",
      "LP"
    ],
    "gold": true
  }
];
```

# --solutions--

```js
const myMusic = [
  {
    "artist": "Billy Joel",
    "title": "Piano Man",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP" ],
    "gold": true
  },
  {
    "artist": "ABBA",
    "title": "Ring Ring",
    "release_year": 1973,
    "formats": [
      "CS",
      "8T",
      "LP",
    "CD",
  ]
  }
];
```
