---
id: 56533eb9ac21ba0edf2244dd
title: switch文で多くの選択肢から選ぶ
challengeType: 1
forumTopicId: 18277
dashedName: selecting-from-many-options-with-switch-statements
---

# --description--

1つの値を多くの選択肢と照合する必要がある場合、<dfn>switch</dfn>文を使うことができます。`switch`文は値をさまざまな可能な値を定義する<dfn>case</dfn>文と比較します。<dfn>case</dfn>ブロック内では有効なJavaScript文を実行でき、最初に一致した`case`の値から`break`が見つかるまで実行されます。

以下は`switch`文の例です：

```js
switch (fruit) {
  case "apple":
    console.log("The fruit is an apple");
    break;
  case "orange":
    console.log("The fruit is an orange");
    break;
}
```

`case`の値は厳密な等価性（`===`）でテストされます。`break`はJavaScriptに文の実行を停止するよう指示します。`break`が省略されると、次の文が実行されます。

# --instructions--

`val`をテストし、次の条件に対して`answer`を設定するswitch文を書いてください：  
`1` - `alpha`  
`2` - `beta`  
`3` - `gamma`  
`4` - `delta`

# --hints--

`caseInSwitch(1)`は文字列`alpha`の値を持つべきです

```js
assert(caseInSwitch(1) === 'alpha');
```

`caseInSwitch(2)`は文字列`beta`の値を持つべきです

```js
assert(caseInSwitch(2) === 'beta');
```

`caseInSwitch(3)`は文字列`gamma`の値を持つべきです

```js
assert(caseInSwitch(3) === 'gamma');
```

`caseInSwitch(4)`は文字列`delta`の値を持つべきです

```js
assert(caseInSwitch(4) === 'delta');
```

`if`や`else`文は使わないでください

```js
assert(!/else/g.test(__helpers.removeJSComments(code)) || !/if/g.test(__helpers.removeJSComments(code)));
```

少なくとも3つの`break`文が必要です

```js
assert(__helpers.removeJSComments(code).match(/break/g).length > 2);
```

# --seed--

## --seed-contents--

```js
function caseInSwitch(val) {
  let answer = "";
  // Only change code below this line



  // Only change code above this line
  return answer;
}

caseInSwitch(1);
```

# --solutions--

```js
function caseInSwitch(val) {
  let answer = "";

  switch (val) {
    case 1:
      answer = "alpha";
      break;
    case 2:
      answer = "beta";
      break;
    case 3:
      answer = "gamma";
      break;
    case 4:
      answer = "delta";
  }
  return answer;
}
```
