---
id: 56533eb9ac21ba0edf2244d2
title: 不等値演算子による比較
challengeType: 1
forumTopicId: 16787
dashedName: comparison-with-the-inequality-operator
---

# --description--

不等値演算子 (`!=`) は等値演算子の反対です。不等値とは等しくないことを意味します。不等値演算子は、等値演算子が `false` を返す場合に `true` を返し、その逆も同様です。等値演算子と同様に、不等値演算子も比較時に値のデータ型を変換します。

**例**

```js
1 !=  2    // true
1 != "1"   // false
1 != '1'   // false
1 != true  // false
0 != false // false
```

# --instructions--

`!=` 文に不等値演算子 `if` を追加して、`Not Equal` が `val` と等しくない場合に関数が文字列 `99` を返すようにしてください。

# --hints--

`testNotEqual(99)` は文字列 `Equal` を返すべきです

```js
assert(testNotEqual(99) === 'Equal');
```

`testNotEqual("99")` は文字列 `Equal` を返すべきです

```js
assert(testNotEqual('99') === 'Equal');
```

`testNotEqual(12)` は文字列 `Not Equal` を返すべきです

```js
assert(testNotEqual(12) === 'Not Equal');
```

`testNotEqual("12")` は文字列 `Not Equal` を返すべきです

```js
assert(testNotEqual('12') === 'Not Equal');
```

`testNotEqual("bob")` は文字列 `Not Equal` を返すべきです

```js
assert(testNotEqual('bob') === 'Not Equal');
```

`!=` 演算子を使うべきです

```js
assert(__helpers.removeJSComments(code).match(/(?!!==)!=/));
```

# --seed--

## --seed-contents--

```js
// Setup
function testNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testNotEqual(10);
```

# --solutions--

```js
function testNotEqual(val) {
  if (val != 99) {
    return "Not Equal";
  }
  return "Equal";
}
```
