---
id: 56533eb9ac21ba0edf2244c4
title: 関数の早期リターンパターン
challengeType: 1
forumTopicId: 18272
dashedName: return-early-pattern-for-functions
---

# --description--

`return` 文に到達すると、現在の関数の実行がストップし、制御が呼び出し元に戻ります。

**例**

```js
function myFun() {
  console.log("Hello");
  return "World";
  console.log("byebye")
}
myFun();
```

上記はコンソールに文字列 `Hello` を表示し、文字列 `World` を返します。文字列 `byebye` はコンソールに表示されません。なぜなら `return` 文で関数が終了するからです。

# --instructions--

関数 `abTest` を修正して、`a` または `b` が `0` より小さい場合、関数がすぐに `undefined` の値で終了するようにしてください。

**ヒント**  
<a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/basic-javascript/understanding-uninitialized-variables" target="_blank" rel="noopener noreferrer nofollow"><code>undefined</code> は文字列ではなくキーワード</a>であることを覚えておいてください。

# --hints--

`abTest(2, 2)` は数値を返すべきです

```js
assert(typeof abTest(2, 2) === 'number');
```

`abTest(2, 2)` は `8` を返すべきです

```js
assert(abTest(2, 2) === 8);
```

`abTest(-2, 2)` は `undefined` を返すべきです

```js
assert(abTest(-2, 2) === undefined);
```

`abTest(2, -2)` は `undefined` を返すべきです

```js
assert(abTest(2, -2) === undefined);
```

`abTest(2, 8)` は `18` を返すべきです

```js
assert(abTest(2, 8) === 18);
```

`abTest(3, 3)` は `12` を返すべきです

```js
assert(abTest(3, 3) === 12);
```

`abTest(0, 0)` は `0` を返すべきです

```js
assert(abTest(0, 0) === 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function abTest(a, b) {
  // Only change code below this line



  // Only change code above this line

  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}

abTest(2,2);
```

# --solutions--

```js
function abTest(a, b) {
  if(a < 0 || b < 0) {
    return undefined;
  }
  return Math.round(Math.pow(Math.sqrt(a) + Math.sqrt(b), 2));
}
```
