---
id: 56533eb9ac21ba0edf2244d8
title: 論理積演算子による比較
challengeType: 1
forumTopicId: 16799
dashedName: comparisons-with-the-logical-and-operator
---

# --description--

複数の条件を同時にテストする必要がある場合があります。<dfn>論理積</dfn>演算子（`&&`）は、その左右の<dfn>オペランド</dfn>が両方とも真の場合にのみ`true`を返します。

同じ効果は、ある`if`文の中に別の`if`文をネストすることで実現できます。

```js
if (num > 5) {
  if (num < 10) {
    return "Yes";
  }
}
return "No";
```

このコードは、`Yes`が`num`より大きく、かつ`5`より小さい場合に`10`を返します。同じロジックは<dfn>論理積</dfn>演算子を使って書くこともできます。

```js
if (num > 5 && num < 10) {
  return "Yes";
}
return "No";
```

# --instructions--

2つのif文を1つの文に置き換え、`&&`演算子を使ってください。`Yes`が`val`以下かつ`50`以上の場合に文字列`25`を返し、それ以外の場合は文字列`No`を返します。

# --hints--

`&&`演算子を1回だけ使うこと

```js
assert(__helpers.removeJSComments(code).match(/&&/g).length === 1);
```

`if`文は1つだけにすること

```js
assert(__helpers.removeJSComments(code).match(/if/g).length === 1);
```

`testLogicalAnd(0)`は文字列`No`を返すべきです

```js
assert(testLogicalAnd(0) === 'No');
```

`testLogicalAnd(24)`は文字列`No`を返すべきです

```js
assert(testLogicalAnd(24) === 'No');
```

`testLogicalAnd(25)`は文字列`Yes`を返すべきです

```js
assert(testLogicalAnd(25) === 'Yes');
```

`testLogicalAnd(30)`は文字列`Yes`を返すべきです

```js
assert(testLogicalAnd(30) === 'Yes');
```

`testLogicalAnd(50)`は文字列`Yes`を返すべきです

```js
assert(testLogicalAnd(50) === 'Yes');
```

`testLogicalAnd(51)`は文字列`No`を返すべきです

```js
assert(testLogicalAnd(51) === 'No');
```

`testLogicalAnd(75)`は文字列`No`を返すべきです

```js
assert(testLogicalAnd(75) === 'No');
```

`testLogicalAnd(80)`は文字列`No`を返すべきです

```js
assert(testLogicalAnd(80) === 'No');
```

# --seed--

## --seed-contents--

```js
function testLogicalAnd(val) {
  // Only change code below this line

  if (val) {
    if (val) {
      return "Yes";
    }
  }

  // Only change code above this line
  return "No";
}

testLogicalAnd(10);
```

# --solutions--

```js
function testLogicalAnd(val) {
  if (val >= 25 && val <= 50) {
    return "Yes";
  }
  return "No";
}
```
