---
id: 56533eb9ac21ba0edf2244be
title: グローバルスコープと関数
challengeType: 1
forumTopicId: 18193
dashedName: global-scope-and-functions
---

# --description--

JavaScriptでは、<dfn>スコープ</dfn>は変数の可視性を指します。関数ブロックの外で定義された変数は<dfn>グローバル</dfn>スコープを持ちます。これは、JavaScriptコードのどこからでも見えることを意味します。

`let`または`const`キーワードなしで宣言された変数は、自動的に`global`スコープで作成されます。これはコードの他の部分や関数を再度実行したときに意図しない結果を生む場合があります。変数は常に`let`または`const`で宣言することをお勧めします。

# --instructions--

`let`または`const`を使って、どの関数の外でもグローバル変数`myGlobal`を宣言してください。値は`10`で初期化してください。

関数`fun1`の中で、`5`、`oopsGlobal`、`var`のキーワードを使わずに、`let`に`const`を代入してください。

# --before-each--

```js
var oopsGlobal;
```

# --hints--

`myGlobal`は定義されている必要があります

```js
assert(typeof myGlobal != 'undefined');
```

`myGlobal`は`10`の値を持つ必要があります

```js
assert(myGlobal === 10);
```

`myGlobal`は`let`または`const`キーワードを使って宣言されている必要があります

```js
assert(/(let|const)\s+myGlobal/.test(__helpers.removeJSComments(code)));
```

`oopsGlobal`はグローバル変数であり、`5`の値を持つ必要があります

```js
fun1();
assert(typeof oopsGlobal != 'undefined');
```

# --seed--

## --seed-contents--

```js
// Declare the myGlobal variable below this line


function fun1() {
  // Assign 5 to oopsGlobal here

}

// Only change code above this line

function fun2() {
  let output = "";
  if (typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if (typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```

# --solutions--

```js
const myGlobal = 10;

function fun1() {
  oopsGlobal = 5;
}

function fun2() {
  let output = "";
  if(typeof myGlobal != "undefined") {
    output += "myGlobal: " + myGlobal;
  }
  if(typeof oopsGlobal != "undefined") {
    output += " oopsGlobal: " + oopsGlobal;
  }
  console.log(output);
}
```
