---
id: 56533eb9ac21ba0edf2244b5
title: 文字列内のリテラル引用符のエスケープ
challengeType: 1
forumTopicId: 17568
dashedName: escaping-literal-quotes-in-strings
---

# --description--

文字列を定義するときは、シングルクオートかダブルクオートで始めて終わらせる必要があります。文字列の中にリテラルの引用符 `"` や `'` を入れたい場合はどうすればよいでしょうか？

JavaScriptでは、引用符の前に <dfn>バックスラッシュ</dfn> (`\`) を置くことで、その引用符を文字列の終わりとみなさないように <dfn>エスケープ</dfn> できます。

```js
const sampleStr = "Alan said, \"Peter is learning JavaScript\".";
```

これにより、JavaScriptは次の引用符が文字列の終わりではなく、文字列の中に表示されるべきだと認識します。たとえばこれをコンソールに出力すると、次のようになります。

```js
Alan said, "Peter is learning JavaScript".
```

# --instructions--

<dfn>バックスラッシュ</dfn>を使って、`myStr` 変数に文字列を代入してください。コンソールに出力すると、次のように表示されるはずです。

```js
I am a "double quoted" string inside "double quotes".
```

# --hints--

ダブルクオート（`"`）を2つ使い、エスケープされたダブルクオート（`\"`）を4つ使う必要があります。

```js
assert(__helpers.removeJSComments(code).match(/\\"/g).length === 4 && __helpers.removeJSComments(code).match(/[^\\]"/g).length === 2);
```

変数 `myStr` には文字列 `I am a "double quoted" string inside "double quotes".` が含まれているべきです。

```js
assert(/I am a "double quoted" string inside "double quotes(\."|"\.)$/.test(myStr));
```

# --seed--

## --seed-contents--

```js
const myStr = ""; // Change this line
```

# --solutions--

```js
const myStr = "I am a \"double quoted\" string inside \"double quotes\".";
```
