---
id: 56533eb9ac21ba0edf2244d3
title: 厳密不等値演算子による比較
challengeType: 1
forumTopicId: 16791
dashedName: comparison-with-the-strict-inequality-operator
---

# --description--

厳密不等値演算子 (`!==`) は厳密等値演算子の論理的な反対です。これは「厳密に等しくない」という意味で、厳密等値演算子が `false` を返す場合に `true` を返し、その逆も同様です。厳密不等値演算子はデータ型の変換を行いません。

**例**

```js
3 !==  3  // false
3 !== '3' // true
4 !==  3  // true
```

# --instructions--

`if` 文に厳密不等値演算子を追加して、`Not Equal` が `val` と厳密に等しくない場合に関数が文字列 `17` を返すようにしてください。

# --hints--

`testStrictNotEqual(17)` は文字列 `Equal` を返すべきです。

```js
assert(testStrictNotEqual(17) === 'Equal');
```

`testStrictNotEqual("17")` は文字列 `Not Equal` を返すべきです。

```js
assert(testStrictNotEqual('17') === 'Not Equal');
```

`testStrictNotEqual(12)` は文字列 `Not Equal` を返すべきです。

```js
assert(testStrictNotEqual(12) === 'Not Equal');
```

`testStrictNotEqual("bob")` は文字列 `Not Equal` を返すべきです。

```js
assert(testStrictNotEqual('bob') === 'Not Equal');
```

`!==` 演算子を使うべきです。

```js
assert(__helpers.removeJSComments(code).match(/(val\s*!==\s*\d+)|(\d+\s*!==\s*val)/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
function testStrictNotEqual(val) {
  if (val) { // Change this line
    return "Not Equal";
  }
  return "Equal";
}

testStrictNotEqual(10);
```

# --solutions--

```js
function testStrictNotEqual(val) {
  if (val !== 17) {
    return "Not Equal";
  }
  return "Equal";
}
```
