---
id: 56533eb9ac21ba0edf2244d6
title: 小なり演算子による比較
challengeType: 1
forumTopicId: 16789
dashedName: comparison-with-the-less-than-operator
---

# --description--

小なり演算子 (`<`) は2つの数値の値を比較します。左側の数が右側の数より小さい場合、`true` を返します。そうでなければ、`false` を返します。等値演算子と同様に、小なり演算子は比較時にデータ型を変換します。

**例**

```js
2   < 5 // true
'3' < 7 // true
5   < 5 // false
3   < 2 // false
'8' < 4 // false
```

# --instructions--

return文が意味をなすように、指定された行に小なり演算子を追加してください。

# --hints--

`testLessThan(0)` は文字列 `Under 25` を返すべきです

```js
assert(testLessThan(0) === 'Under 25');
```

`testLessThan(24)` は文字列 `Under 25` を返すべきです

```js
assert(testLessThan(24) === 'Under 25');
```

`testLessThan(25)` は文字列 `Under 55` を返すべきです

```js
assert(testLessThan(25) === 'Under 55');
```

`testLessThan(54)` は文字列 `Under 55` を返すべきです

```js
assert(testLessThan(54) === 'Under 55');
```

`testLessThan(55)` は文字列 `55 or Over` を返すべきです

```js
assert(testLessThan(55) === '55 or Over');
```

`testLessThan(99)` は文字列 `55 or Over` を返すべきです

```js
assert(testLessThan(99) === '55 or Over');
```

小なり演算子 `<` を少なくとも2回使う必要があります

```js
assert(__helpers.removeJSComments(code).match(/val\s*<\s*('|")*\d+('|")*/g).length > 1);
```

# --seed--

## --seed-contents--

```js
function testLessThan(val) {
  if (val) {  // Change this line
    return "Under 25";
  }

  if (val) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}

testLessThan(10);
```

# --solutions--

```js
function testLessThan(val) {
  if (val < 25) {  // Change this line
    return "Under 25";
  }

  if (val < 55) {  // Change this line
    return "Under 55";
  }

  return "55 or Over";
}
```
