---
id: 56592a60ddddeae28f7aa8e1
title: インデックスで多次元配列にアクセスする
challengeType: 1
forumTopicId: 16159
dashedName: access-multi-dimensional-arrays-with-indexes
---

# --description--

<dfn>多次元</dfn>配列を考える一つの方法は、*配列の配列*としてイメージすることです。配列にアクセスするとき、最初の括弧のセットは最も外側（第一レベル）の配列のエントリーを指し、追加の括弧のペアは内側の次のレベルのエントリーを指します。

**例**

```js
const arr = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14]
];

const subarray = arr[3];
const nestedSubarray = arr[3][0];
const element = arr[3][0][1];
```

この例では、`subarray` は `[[10, 11, 12], 13, 14]` の値を持ち、
`nestedSubarray` は `[10, 11, 12]` の値を持ち、`element` は `11` の値を持っています。

**注意:** 配列名と大括弧の間にスペースを入れてはいけません。例えば `array [0][0]` や `array [0] [0]` のような書き方は許されません。JavaScriptはこれを正しく処理できますが、他のプログラマーがコードを読むときに混乱を招く可能性があります。

# --instructions--

ブラケット記法を使って `myArray` から要素を選び、`myData` が `8` と等しくなるようにしてください。

# --hints--

`myData` は `8` と等しくなるはずです。

```js
assert(myData === 8);
```

`myArray` から正しい値を読み取るためにブラケット記法を使う必要があります。

```js
assert(/myData=myArray\[2\]\[1\]/.test(__helpers.removeWhiteSpace(__helpers.removeJSComments(code))));
```

# --seed--

## --seed-contents--

```js
const myArray = [
  [1, 2, 3],
  [4, 5, 6],
  [7, 8, 9],
  [[10, 11, 12], 13, 14],
];

const myData = myArray[0][0];
```

# --solutions--

```js
const myArray = [[1, 2, 3], [4, 5, 6], [7, 8, 9], [[10, 11, 12], 13, 14]];
const myData = myArray[2][1];
```
