---
id: 565bbe00e9cc8ac0725390f4
title: カードカウンティング
challengeType: 1
forumTopicId: 16809
dashedName: counting-cards
---

# --description--

カジノゲームのブラックジャックでは、プレイヤーはデッキに残っている高いカードと低いカードの相対的な数を追跡することで、次の手でハウスに対して有利かどうかを判断できます。これをカードカウンティングと呼びます。

デッキに高いカードが多く残っているとプレイヤーに有利です。各カードには下のテーブルに従って値が代入されます。カウントが正の値のときはプレイヤーは高額ベットをすべきです。カウントがゼロまたは負の値のときは低額ベットをすべきです。

<table><thead><tr><th>カウント変化</th><th>カード</th></tr></thead><tbody><tr><td>+1</td><td>2、3、4、5、6</td></tr><tr><td>0</td><td>7、8、9</td></tr><tr><td>-1</td><td>10、'J'、'Q'、'K'、'A'</td></tr></tbody></table>

カードカウンティングの関数を作成してください。この関数は`card`というパラメータを受け取ります。これは数字か文字列で、カードの値に応じてグローバル変数`count`を増減させます（テーブル参照）。その後、現在のカウントと、カウントが正のときは`Bet`、ゼロまたは負のときは`Hold`という文字列を返します。現在のカウントとプレイヤーの判断（`Bet`または`Hold`）は半角スペース1つで区切ってください。

**出力例:** `-3 Hold` または `5 Bet`

**ヒント**  
値が7、8、9のときに`count`を0にリセットしないでください。  
配列を返さないでください。  
出力にシングルクォートやダブルクォートを含めないでください。

# --hints--

関数はカウントの値とテキスト（`Bet` または `Hold`）を半角スペース1つで返す必要があります。

```js
assert(//
  (function () {
    count = 0;
    let out = cc(10);
    const hasSpace = /-?\d+ (Bet|Hold)/.test('' + out);
    return hasSpace;
  })()
);
```

カードのシーケンス 2、3、4、5、6 は文字列 `5 Bet` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(3);
    cc(4);
    cc(5);
    var out = cc(6);
    if (out === '5 Bet') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 7、8、9 は文字列 `0 Hold` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(7);
    cc(8);
    var out = cc(9);
    if (out === '0 Hold') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 10、J、Q、K、A は文字列 `-5 Hold` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(10);
    cc('J');
    cc('Q');
    cc('K');
    var out = cc('A');
    if (out === '-5 Hold') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 3、7、Q、8、A は文字列 `-1 Hold` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(7);
    cc('Q');
    cc(8);
    var out = cc('A');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 2、J、9、2、7 は文字列 `1 Bet` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc('J');
    cc(9);
    cc(2);
    var out = cc(7);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 2、2、10 は文字列 `1 Bet` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(2);
    cc(2);
    var out = cc(10);
    if (out === '1 Bet') {
      return true;
    }
    return false;
  })()
);
```

カードのシーケンス 3、2、A、10、K は文字列 `-1 Hold` を返すべきです。

```js
assert(
  (function () {
    count = 0;
    cc(3);
    cc(2);
    cc('A');
    cc(10);
    var out = cc('K');
    if (out === '-1 Hold') {
      return true;
    }
    return false;
  })()
);
```

# --seed--

## --seed-contents--

```js
let count = 0;

function cc(card) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

cc(2); cc(3); cc(7); cc('K'); cc('A');
```

# --solutions--

```js
let count = 0;
function cc(card) {
  switch(card) {
    case 2:
    case 3:
    case 4:
    case 5:
    case 6:
      count++;
      break;
    case 10:
    case 'J':
    case 'Q':
    case 'K':
    case 'A':
      count--;
  }
  if(count > 0) {
    return count + " Bet";
  } else {
    return count + " Hold";
  }
}
```
