---
id: 5664820f61c48e80c9fa476c
title: ゴルフコード
challengeType: 1
forumTopicId: 18195
dashedName: golf-code
---

# --description--

ゴルフのゲームでは、各ホールに`par`があり、これはゴルファーがボールをホールに入れてプレイを完了するために期待される平均`strokes`の数を意味します。`par`からあなたの`strokes`がどれだけ上回っているか下回っているかによって、異なるニックネームがあります。

あなたの関数には`par`と`strokes`の引数が渡されます。優先度の高い順（上から下）に`strokes`をリストしたこのテーブルに従って、正しい文字列を返してください。

<table><thead><tr><th>ストローク</th><th>返す文字列</th></tr></thead><tbody><tr><td>1</td><td>"Hole-in-one!"</td></tr><tr><td>&#x3C;= par - 2</td><td>"Eagle"</td></tr><tr><td>par - 1</td><td>"Birdie"</td></tr><tr><td>par</td><td>"Par"</td></tr><tr><td>par + 1</td><td>"Bogey"</td></tr><tr><td>par + 2</td><td>"Double Bogey"</td></tr><tr><td>>= par + 3</td><td>"Go Home!"</td></tr></tbody></table>

`par`と`strokes`は常に数値で正の値です。利便性のためにすべての名前を配列にまとめています。

# --hints--

`golfScore(4, 1)`は文字列`Hole-in-one!`を返すべきです

```js
assert(golfScore(4, 1) === 'Hole-in-one!');
```

`golfScore(4, 2)`は文字列`Eagle`を返すべきです

```js
assert(golfScore(4, 2) === 'Eagle');
```

`golfScore(5, 2)`は文字列`Eagle`を返すべきです

```js
assert(golfScore(5, 2) === 'Eagle');
```

`golfScore(4, 3)`は文字列`Birdie`を返すべきです

```js
assert(golfScore(4, 3) === 'Birdie');
```

`golfScore(4, 4)`は文字列`Par`を返すべきです

```js
assert(golfScore(4, 4) === 'Par');
```

`golfScore(1, 1)`は文字列`Hole-in-one!`を返すべきです

```js
assert(golfScore(1, 1) === 'Hole-in-one!');
```

`golfScore(5, 5)`は文字列`Par`を返すべきです

```js
assert(golfScore(5, 5) === 'Par');
```

`golfScore(4, 5)`は文字列`Bogey`を返すべきです

```js
assert(golfScore(4, 5) === 'Bogey');
```

`golfScore(4, 6)`は文字列`Double Bogey`を返すべきです

```js
assert(golfScore(4, 6) === 'Double Bogey');
```

`golfScore(4, 7)`は文字列`Go Home!`を返すべきです

```js
assert(golfScore(4, 7) === 'Go Home!');
```

`golfScore(5, 9)`は文字列`Go Home!`を返すべきです

```js
assert(golfScore(5, 9) === 'Go Home!');
```

# --seed--

## --seed-contents--

```js
const names = ["Hole-in-one!", "Eagle", "Birdie", "Par", "Bogey", "Double Bogey", "Go Home!"];

function golfScore(par, strokes) {
  // Only change code below this line


  return "Change Me";
  // Only change code above this line
}

golfScore(5, 4);
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) {
    return "Hole-in-one!";
  }

  if (strokes <= par - 2) {
    return "Eagle";
  }

  if (strokes === par - 1) {
    return "Birdie";
  }

  if (strokes === par) {
    return "Par";
  }

  if (strokes === par + 1) {
    return "Bogey";
  }

  if(strokes === par + 2) {
    return "Double Bogey";
  }

  return "Go Home!";
}
```
