---
id: 5679ceb97cbaa8c51670a16b
title: 関数からブール値を返す
challengeType: 1
forumTopicId: 18273
dashedName: returning-boolean-values-from-functions
---

# --description--

<a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/comparison-with-the-equality-operator" target="_blank" rel="noopener noreferrer nofollow">等値演算子による比較</a>で、すべての比較演算子はブール値の`true`または`false`を返すことを覚えているかもしれません。

時々、次のように`if/else`文を使って比較を行う人がいます。

```js
function isEqual(a, b) {
  if (a === b) {
    return true;
  } else {
    return false;
  }
}
```

しかし、これより良い方法があります。`===`は`true`または`false`を返すので、比較の結果をそのまま返すことができます。

```js
function isEqual(a, b) {
  return a === b;
}
```

# --instructions--

`isLess`関数を修正して、`if/else`文を取り除いてください。

# --hints--

`isLess(10, 15)`は`true`を返すべきです

```js
assert(isLess(10, 15) === true);
```

`isLess(15, 10)`は`false`を返すべきです

```js
assert(isLess(15, 10) === false);
```

`if`や`else`文は使わないでください

```js
assert(!/if|else/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function isLess(a, b) {
  // Only change code below this line
  if (a < b) {
    return true;
  } else {
    return false;
  }
  // Only change code above this line
}

isLess(10, 15);
```

# --solutions--

```js
function isLess(a, b) {
  return a < b;
}
```
