---
id: 5688e62ea601b2482ff8422b
title: プロフィール検索
challengeType: 1
forumTopicId: 18259
dashedName: profile-lookup
---

# --description--

私たちの連絡先リストにあるさまざまな人を表すオブジェクトの配列があります。

`lookUpProfile` 関数は `name` とプロパティ（`prop`）を引数として受け取るようにあらかじめ作成されています。

この関数は、`name` が実際の連絡先の `firstName` であり、指定されたプロパティ（`prop`）がその連絡先のプロパティであるかどうかを確認する必要があります。

両方が真であれば、そのプロパティの「値」を返してください。

`name` がどの連絡先にも対応しない場合は、文字列 `No such contact` を返してください。

`prop` が `name` に一致する連絡先の有効なプロパティに対応しない場合は、文字列 `No such property` を返してください。

# --hints--

`lookUpProfile("Kristian", "lastName")` は文字列 `Vos` を返すべきです

```js
assert(lookUpProfile('Kristian', 'lastName') === 'Vos');
```

`lookUpProfile("Sherlock", "likes")` は `["Intriguing Cases", "Violin"]` を返すべきです

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

`lookUpProfile("Harry", "likes")` は配列を返すべきです

```js
assert(typeof lookUpProfile('Harry', 'likes') === 'object');
```

`lookUpProfile("Bob", "number")` は文字列 `No such contact` を返すべきです

```js
assert(lookUpProfile('Bob', 'number') === 'No such contact');
```

`lookUpProfile("Bob", "potato")` は文字列 `No such contact` を返すべきです

```js
assert(lookUpProfile('Bob', 'potato') === 'No such contact');
```

`lookUpProfile("Akira", "address")` は文字列 `No such property` を返すべきです

```js
assert(lookUpProfile('Akira', 'address') === 'No such property');
```

# --seed--

## --seed-contents--

```js
// Setup
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

function lookUpProfile(name, prop) {
  // Only change code below this line

  // Only change code above this line
}

lookUpProfile("Akira", "likes");
```

# --solutions--

```js
const contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];
function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
