---
id: 5690307fddb111c6084545d7
title: "`if` `else if` 文の論理的な順序"
challengeType: 1
forumTopicId: 18228
dashedName: logical-order-in-if-else-statements
---

# --description--

`if`、`else if` 文では順序が重要です。

関数は上から下へ実行されるため、どの文が最初に来るかに注意する必要があります。

この2つの関数を例に見てみましょう。

最初のものはこちらです：

```js
function foo(x) {
  if (x < 1) {
    return "Less than one";
  } else if (x < 2) {
    return "Less than two";
  } else {
    return "Greater than or equal to two";
  }
}
```

そして2つ目は文の順序を入れ替えただけです：

```js
function bar(x) {
  if (x < 2) {
    return "Less than two";
  } else if (x < 1) {
    return "Less than one";
  } else {
    return "Greater than or equal to two";
  }
}
```

これら2つの関数はほとんど同じに見えますが、どちらにも数値を渡すと異なる出力が得られます。

```js
foo(0)
bar(0)
```

`foo(0)` は文字列 `Less than one` を返し、`bar(0)` は文字列 `Less than two` を返します。

# --instructions--

すべての場合で正しい文を返すように、関数内の論理の順序を変更してください。

# --hints--

`orderMyLogic(4)` は文字列 `Less than 5` を返すべきです

```js
assert(orderMyLogic(4) === 'Less than 5');
```

`orderMyLogic(6)` は文字列 `Less than 10` を返すべきです

```js
assert(orderMyLogic(6) === 'Less than 10');
```

`orderMyLogic(11)` は文字列 `Greater than or equal to 10` を返すべきです

```js
assert(orderMyLogic(11) === 'Greater than or equal to 10');
```

# --seed--

## --seed-contents--

```js
function orderMyLogic(val) {
  if (val < 10) {
    return "Less than 10";
  } else if (val < 5) {
    return "Less than 5";
  } else {
    return "Greater than or equal to 10";
  }
}

orderMyLogic(7);
```

# --solutions--

```js
function orderMyLogic(val) {
  if(val < 5) {
    return "Less than 5";
  } else if (val < 10) {
    return "Less than 10";
  } else {
    return "Greater than or equal to 10";
  }
}
```
