---
id: 56bbb991ad1ed5201cd392d0
title: JavaScriptオブジェクトを作成する
challengeType: 1
forumTopicId: 16769
dashedName: build-javascript-objects
---

# --description--

`object` という用語を聞いたことがあるかもしれません。

オブジェクトは `arrays` に似ていますが、データにアクセスしたり変更したりするのにインデックスを使う代わりに、オブジェクトでは `properties` と呼ばれるものを通じてデータにアクセスします。

オブジェクトはデータを構造化して保存するのに便利で、猫のような現実世界のオブジェクトを表すこともできます。

以下は猫のオブジェクトの例です：

```js
const cat = {
  "name": "Whiskers",
  "legs": 4,
  "tails": 1,
  "enemies": ["Water", "Dogs"]
};
```

この例では、すべてのプロパティが `name`、`legs`、`tails` のように文字列として保存されています。しかし、プロパティに数値を使うこともできます。単語1つの文字列プロパティの場合は、引用符を省略することも可能です：

```js
const anotherObject = {
  make: "Ford",
  5: "five",
  "model": "focus"
};
```

ただし、オブジェクトに文字列以外のプロパティがある場合、JavaScriptはそれらを自動的に文字列に型変換します。

# --instructions--

`myDog` という名前の犬を表すオブジェクトを作成してください。このオブジェクトには `name`（文字列）、`legs`、`tails`、`friends` のプロパティを含めてください。

これらのオブジェクトのプロパティには好きな値を設定できますが、`name` は文字列、`legs` と `tails` は数値、`friends` は配列である必要があります。

# --hints--

`myDog` は `name` というプロパティを含み、それは `string` であるべきです。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('name') &&
      z.name !== undefined &&
      typeof z.name === 'string'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` は `legs` というプロパティを含み、それは `number` であるべきです。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('legs') &&
      z.legs !== undefined &&
      typeof z.legs === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` は `tails` というプロパティを含み、それは `number` であるべきです。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('tails') &&
      z.tails !== undefined &&
      typeof z.tails === 'number'
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` は `friends` というプロパティを含み、それは `array` であるべきです。

```js
assert(
  (function (z) {
    if (
      z.hasOwnProperty('friends') &&
      z.friends !== undefined &&
      Array.isArray(z.friends)
    ) {
      return true;
    } else {
      return false;
    }
  })(myDog)
);
```

`myDog` は与えられたすべてのプロパティのみを含むべきです。

```js
assert(
  (function (z) {
    return Object.keys(z).length === 4;
  })(myDog)
);
```

# --seed--

## --seed-contents--

```js
const myDog = {
  // Only change code below this line


  // Only change code above this line
};
```

# --solutions--

```js
const myDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```
