---
id: 56bbb991ad1ed5201cd392d3
title: JavaScriptオブジェクトからプロパティを削除する
challengeType: 1
forumTopicId: 17560
dashedName: delete-properties-from-a-javascript-object
---

# --description--

オブジェクトからプロパティを次のように削除することもできます。

```js
delete ourDog.bark;
```

例：

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"],
  "bark": "bow-wow"
};

delete ourDog.bark;
```

上記の最後の行の後、`ourDog` は次のようになります。

```js
{
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
}
```

# --instructions--

`tails`から`myDog`プロパティを削除してください。ドット記法かブラケット記法のどちらかを使えます。

# --hints--

`tails`から`myDog`プロパティを削除する必要があります。

```js
assert(typeof myDog === 'object' && myDog.tails === undefined);
```

`myDog`の設定は変更しないでください。

```js
assert(__helpers.removeJSComments(code).match(/"tails": 1/g).length > 0);
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"],
  "bark": "woof"
};
delete myDog.tails;
```
