---
id: 56bbb991ad1ed5201cd392ce
title: "`unshift` メソッドで配列を操作する"
challengeType: 1
forumTopicId: 18239
dashedName: manipulate-arrays-with-unshift
---

# --description--

配列の先頭から要素を `shift` できるだけでなく、配列の先頭に要素を `unshift` することもできます。つまり、配列の前に要素を追加できます。

`.unshift()` は `.push()` とまったく同じように動作しますが、配列の末尾に要素を追加する代わりに、`unshift()` は配列の先頭に要素を追加します。

例：

```js
const ourArray = ["Stimpson", "J", "cat"];
ourArray.shift();
ourArray.unshift("Happy");
```

`shift` の後、`ourArray` は `["J", "cat"]` の値を持ちます。`unshift` の後、`ourArray` は `["Happy", "J", "cat"]` の値を持ちます。

# --instructions--

`unshift()` を使って、`myArray` 変数の先頭に `["Paul", 35]` を追加してください。

# --hints--

`myArray` は今、`[["Paul", 35], ["dog", 3]]` を持つはずです。

```js
assert(
  (function (d) {
    if (
      typeof d[0] === 'object' &&
      d[0][0] == 'Paul' &&
      d[0][1] === 35 &&
      d[1][0] != undefined &&
      d[1][0] == 'dog' &&
      d[1][1] != undefined &&
      d[1][1] == 3
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];
myArray.shift();
myArray.unshift(["Paul", 35]);
```
