---
id: 56bbb991ad1ed5201cd392d1
title: オブジェクトのプロパティを更新する
challengeType: 1
forumTopicId: 18336
dashedName: updating-object-properties
---

# --description--

JavaScriptのオブジェクトを作成した後は、他の変数を更新するのと同じようにいつでもそのプロパティを更新できます。更新にはドット記法かブラケット記法のどちらかを使えます。

例えば、`ourDog`を見てみましょう。

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};
```

彼は特に幸せな犬なので、名前を文字列の`Happy Camper`に変えましょう。オブジェクトのnameプロパティを更新するには、`ourDog.name = "Happy Camper";`または`ourDog["name"] = "Happy Camper";`のようにします。これで`ourDog.name`を評価すると、`Camper`ではなく新しい名前の`Happy Camper`が返されます。

# --instructions--

`myDog`オブジェクトのnameプロパティを更新してください。名前を`Coder`から`Happy Coder`に変えましょう。ドット記法かブラケット記法のどちらかを使えます。

# --hints--

`myDog`の`name`プロパティを文字列の`Happy Coder`に更新してください。

```js
assert(/happy coder/gi.test(myDog.name));
```

`myDog`の定義は編集しないでください。

```js
assert(/"name": "Coder"/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};

// Only change code below this line

```

# --solutions--

```js
const myDog = {
  "name": "Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.name = "Happy Coder";
```
