---
id: 56bbb991ad1ed5201cd392cf
title: 関数で再利用可能なJavaScriptを書く
challengeType: 1
forumTopicId: 18378
dashedName: write-reusable-javascript-with-functions
---

# --description--

JavaScriptでは、コードを<dfn>関数</dfn>と呼ばれる再利用可能な部分に分けることができます。

関数の例を示します：

```js
function functionName() {
  console.log("Hello World");
}
```

この関数は名前の後に丸括弧を付けて呼び出すことができます。例えば`functionName();` 関数が呼び出されるたびに、開発者コンソールにメッセージ`Hello World`が出力されます。波括弧の間にあるすべてのコードは、関数が呼び出されるたびに実行されます。

# --instructions--

<ol>
  <li>
    <code>reusableFunction</code>という名前の関数を作成し、開発者コンソールに文字列<code>Hi World</code>を出力してください。
  </li>
  <li>
    その関数を呼び出してください。
  </li>
</ol>

# --before-each--

```js
function testConsole() {
  var originalConsole = console;
  var nativeLog = console.log;
  var hiWorldWasLogged = false;
  console.log = function (message) {
    if (message === 'Hi World') {
      console.warn(message);
      hiWorldWasLogged = true;
    }
    if (nativeLog.apply) {
      nativeLog.apply(originalConsole, arguments);
    } else {
      var nativeMsg = Array.prototype.slice.apply(arguments).join(' ');
      nativeLog(nativeMsg);
    }
  };
  reusableFunction();
  console.log = nativeLog;
  return hiWorldWasLogged;
}
```

# --hints--

`reusableFunction`は関数である必要があります。

```js
assert(typeof reusableFunction === 'function');
```

`reusableFunction`が呼び出されると、文字列`Hi World`がコンソールに出力されるはずです。

```js
assert(testConsole());
```

`reusableFunction`は定義されたらすぐに呼び出してください。

```js
const functionStr = reusableFunction && __helpers.removeWhiteSpace(reusableFunction.toString());
const codeWithoutFunction = __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).replace(/reusableFunction\(\)\{/g, '');
assert(/reusableFunction\(\)/.test(codeWithoutFunction));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function reusableFunction() {
  console.log("Hi World");
}
reusableFunction();
```
