---
id: 56bbb991ad1ed5201cd392ca
title: インデックスで配列データにアクセスする
challengeType: 1
forumTopicId: 16158
dashedName: access-array-data-with-indexes
---

# --description--

配列の中のデータには<dfn>インデックス</dfn>を使ってアクセスできます。

配列のインデックスは文字列で使うブラケット記法と同じ書き方ですが、文字を指定する代わりに配列の要素を指定します。文字列と同様に、配列も<dfn>ゼロベース</dfn>のインデックスを使うので、配列の最初の要素はインデックス`0`になります。

<br>

**例**

```js
const array = [50, 60, 70];
console.log(array[0]);
const data = array[1];
```

`console.log(array[0])`は`50`を出力し、`data`の値は`60`です。

# --instructions--

`myData`という変数を作成し、ブラケット記法を使って`myArray`の最初の値と等しくなるように設定してください。

# --hints--

変数`myData`は`myArray`の最初の値と等しくなる必要があります。

```js
assert(
  (function () {
    if (
      typeof myArray !== 'undefined' &&
      typeof myData !== 'undefined' &&
      myArray[0] === myData
    ) {
      return true;
    } else {
      return false;
    }
  })()
);
```

変数`myArray`のデータにはブラケット記法でアクセスしてください。

```js
assert(
  (function () {
    if (__helpers.removeJSComments(code).match(/\s*=\s*myArray\[0\]/g)) {
      return true;
    } else {
      return false;
    }
  })()
);
```

# --seed--

## --seed-contents--

```js
const myArray = [50, 60, 70];


```

# --solutions--

```js
const myArray = [50, 60, 70];
const myData = myArray[0];
```
