---
id: 56bbb991ad1ed5201cd392cd
title: shift メソッドで配列を操作する
challengeType: 1
forumTopicId: 18238
dashedName: manipulate-arrays-with-shift
---

# --description--

`pop()` は常に配列の最後の要素を削除します。では、最初の要素を削除したい場合はどうすればいいでしょうか？

そこで `.shift()` の出番です。`.pop()` と同じように動作しますが、最後の要素ではなく最初の要素を削除します。

例：

```js
const ourArray = ["Stimpson", "J", ["cat"]];
const removedFromOurArray = ourArray.shift();
```

`removedFromOurArray` は文字列 `Stimpson` の値を持ち、`ourArray` は `["J", ["cat"]]` を持ちます。

# --instructions--

`.shift()` 関数を使って `myArray` の最初の項目を削除し、「shift された」値を新しい変数 `removedFromMyArray` に代入してください。

# --hints--

`myArray` は今 `[["dog", 3]]` と等しくなっているはずです。

```js
assert(
  (function (d) {
    if (d[0][0] == 'dog' && d[0][1] === 3 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`removedFromMyArray` は `["John", 23]` を含んでいるはずです。

```js
assert(
  (function (d) {
    if (
      d[0] == 'John' &&
      d[1] === 23 &&
      typeof removedFromMyArray === 'object'
    ) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["dog", 3]];

// Only change code below this line
const removedFromMyArray = myArray.shift();
```
