---
id: 56bbb991ad1ed5201cd392cb
title: push メソッドで配列を操作する
challengeType: 1
forumTopicId: 18237
dashedName: manipulate-arrays-with-push
---

# --description--

配列の末尾にデータを追加する簡単な方法は `push()` メソッドです。

`push()` メソッドは1つ以上の<dfn>引数</dfn>を受け取り、それらを配列の末尾に引数の順番で追加します。戻り値は配列の新しい長さです。

例:

```js
const arr1 = [1, 2, 3];
arr1.push(4, 5);

const arr2 = ["Stimpson", "J", "cat"];
arr2.push(["happy", "joy"]);
```

`arr1` は現在 `[1, 2, 3, 4, 5]` の値を持ち、`arr2` は `["Stimpson", "J", "cat", ["happy", "joy"]]` の値を持っています。

# --instructions--

`["dog", 3]` を `myArray` 変数の末尾に追加してください。

# --hints--

`myArray` は現在 `[["John", 23], ["cat", 2], ["dog", 3]]` と等しくなっているはずです。

```js
assert(
  (function (d) {
    if (
      d[2] != undefined &&
      d[0][0] == 'John' &&
      d[0][1] === 23 &&
      d[2][0] == 'dog' &&
      d[2][1] === 3 &&
      d[2].length == 2
    ) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
myArray.push(["dog",3]);
```
