---
id: 56bbb991ad1ed5201cd392d2
title: JavaScriptオブジェクトに新しいプロパティを追加する
challengeType: 1
forumTopicId: 301169
dashedName: add-new-properties-to-a-javascript-object
---

# --description--

既存のJavaScriptオブジェクトに新しいプロパティを追加する方法は、変更する場合と同じです。

`bark`に`ourDog`プロパティを追加する方法は次の通りです。

```js
ourDog.bark = "bow-wow";
```

または

```js
ourDog["bark"] = "bow-wow";
```

これで`ourDog.bark`を評価すると、彼の鳴き声である`bow-wow`が得られます。

例：

```js
const ourDog = {
  "name": "Camper",
  "legs": 4,
  "tails": 1,
  "friends": ["everything!"]
};

ourDog.bark = "bow-wow";
```

# --instructions--

`bark`に`myDog`プロパティを追加し、「woof」のような犬の鳴き声を設定してください。ドット記法またはブラケット記法のどちらかを使えます。

# --hints--

`bark`に`myDog`プロパティを追加してください。

```js
assert(myDog.bark !== undefined);
```

`bark`の初期化時に`myDog`を追加しないでください。

```js
assert(!/bark[^\n]:/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};


```

# --solutions--

```js
const myDog = {
  "name": "Happy Coder",
  "legs": 4,
  "tails": 1,
  "friends": ["freeCodeCamp Campers"]
};
myDog.bark = "Woof Woof";
```
