---
id: 56bbb991ad1ed5201cd392cc
title: pop メソッドで配列を操作する
challengeType: 1
forumTopicId: 18236
dashedName: manipulate-arrays-with-pop
---

# --description--

配列のデータを変更するもう一つの方法は `.pop()` 関数を使うことです。

`.pop()` は配列の最後から値を取り出すために使います。この取り出した値は変数に代入して保存できます。つまり、`.pop()` は配列の最後の要素を削除し、その要素を返します。

配列から取り出せる値の種類は何でもかまいません。数値、文字列、さらにはネストされた配列も可能です。

```js
const threeArr = [1, 4, 6];
const oneDown = threeArr.pop();
console.log(oneDown);
console.log(threeArr);
```

最初の `console.log` は値 `6` を表示し、2番目は値 `[1, 4]` を表示します。

# --instructions--

`.pop()` 関数を使って `myArray` の最後の項目を削除し、取り出した値を新しい変数 `removedFromMyArray` に代入してください。

# --hints--

`myArray` には `[["John", 23]]` のみを含めるべきです。

```js
assert(
  (function (d) {
    if (d[0][0] == 'John' && d[0][1] === 23 && d[1] == undefined) {
      return true;
    } else {
      return false;
    }
  })(myArray)
);
```

`pop()` を `myArray` に使うべきです。

```js
assert(/removedFromMyArray\s*=\s*myArray\s*.\s*pop\s*(\s*)/.test(__helpers.removeJSComments(code)));
```

`removedFromMyArray` には `["cat", 2]` のみを含めるべきです。

```js
assert(
  (function (d) {
    if (d[0] == 'cat' && d[1] === 2 && d[2] == undefined) {
      return true;
    } else {
      return false;
    }
  })(removedFromMyArray)
);
```

# --seed--

## --seed-contents--

```js
// Setup
const myArray = [["John", 23], ["cat", 2]];

// Only change code below this line

```

# --solutions--

```js
const myArray = [["John", 23], ["cat", 2]];
const removedFromMyArray = myArray.pop();
```
