---
id: 587d7b7e367417b2b2512b23
title: "parseInt 関数の使い方"
challengeType: 1
forumTopicId: 301183
dashedName: use-the-parseint-function
---

# --description--

`parseInt()` 関数は文字列をパースして整数を返します。例を示します：

```js
const a = parseInt("007");
```

上記の関数は文字列 `007` を整数 `7` に変換します。文字列の最初の文字が数字に変換できない場合は、`NaN` を返します。

# --instructions--

`parseInt()` を `convertToInteger` 関数内で使い、入力文字列 `str` を整数に変換して返してください。

# --hints--

`convertToInteger` は `parseInt()` 関数を使うべきです

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("56")` は数値を返すべきです

```js
assert(typeof convertToInteger('56') === 'number');
```

`convertToInteger("56")` は 56 を返すべきです

```js
assert(convertToInteger('56') === 56);
```

`convertToInteger("77")` は 77 を返すべきです

```js
assert(convertToInteger('77') === 77);
```

`convertToInteger("JamesBond")` は `NaN` を返すべきです

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("56");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str);
}
```
