---
id: 587d7b7e367417b2b2512b24
title: 三項演算子を使う
challengeType: 1
forumTopicId: 301181
dashedName: use-the-conditional-ternary-operator
---

# --description--

<dfn>条件演算子</dfn>は、<dfn>三項演算子</dfn>とも呼ばれ、1行で書ける if-else 式として使えます。

構文は `a ? b : c` で、`a` が条件、`b` が条件が `true` のときに実行されるコード、`c` が条件が `false` のときに実行されるコードです。

次の関数は `if/else` 文を使って条件をチェックしています。

```js
function findGreater(a, b) {
  if(a > b) {
    return "a is greater";
  }
  else {
    return "b is greater or equal";
  }
}
```

これを条件演算子を使って書き直すことができます。

```js
function findGreater(a, b) {
  return a > b ? "a is greater" : "b is greater or equal";
}
```

# --instructions--

`checkEqual` 関数で条件演算子を使い、2つの数字が等しいかどうかをチェックしてください。関数は文字列 `Equal` または文字列 `Not Equal` のどちらかを返す必要があります。

# --hints--

`checkEqual` は条件演算子を使うべきです

```js
assert(/.+?\s*?\?\s*?.+?\s*?:\s*?.+?/.test(__helpers.removeJSComments(code)));
```

`checkEqual(1, 2)` は文字列 `Not Equal` を返すべきです

```js
assert(checkEqual(1, 2) === 'Not Equal');
```

`checkEqual(1, 1)` は文字列 `Equal` を返すべきです

```js
assert(checkEqual(1, 1) === 'Equal');
```

`checkEqual(1, -1)` は文字列 `Not Equal` を返すべきです

```js
assert(checkEqual(1, -1) === 'Not Equal');
```

# --seed--

## --seed-contents--

```js
function checkEqual(a, b) {

}

checkEqual(1, 2);
```

# --solutions--

```js
function checkEqual(a, b) {
  return a === b ? "Equal" : "Not Equal";
}
```
