---
id: 587d7b7e367417b2b2512b22
title: 基数を指定して parseInt 関数を使う
challengeType: 1
forumTopicId: 301182
dashedName: use-the-parseint-function-with-a-radix
---

# --description--

`parseInt()` 関数は文字列をパースして整数を返します。第2引数に基数を指定でき、これは文字列内の数値の基数を表します。基数は2から36までの整数です。

関数の呼び出しは次のようになります。

```js
parseInt(string, radix);
```

例を示します。

```js
const a = parseInt("11", 2);
```

radix 変数は `11` が2進数、つまり基数2であることを示しています。この例では文字列 `11` を整数 `3` に変換しています。

# --instructions--

`parseInt()` を `convertToInteger` 関数内で使い、2進数の数値を整数に変換して返してください。

# --hints--

`convertToInteger` は `parseInt()` 関数を使うべきです

```js
assert(/parseInt/g.test(__helpers.removeJSComments(code)));
```

`convertToInteger("10011")` は数値を返すべきです

```js
assert(typeof convertToInteger('10011') === 'number');
```

`convertToInteger("10011")` は 19 を返すべきです

```js
assert(convertToInteger('10011') === 19);
```

`convertToInteger("111001")` は 57 を返すべきです

```js
assert(convertToInteger('111001') === 57);
```

`convertToInteger("JamesBond")` は `NaN` を返すべきです

```js
assert.isNaN(convertToInteger('JamesBond'));
```

# --seed--

## --seed-contents--

```js
function convertToInteger(str) {

}

convertToInteger("10011");
```

# --solutions--

```js
function convertToInteger(str) {
  return parseInt(str, 2);
}
```
