---
id: 587d7b87367417b2b2512b3f
title: var と let キーワードの違いを探る
challengeType: 1
forumTopicId: 301202
dashedName: explore-differences-between-the-var-and-let-keywords
---

# --description--

`var` キーワードで変数を宣言する際の最大の問題の1つは、変数宣言を簡単に上書きできてしまうことです。

```js
var camper = "James";
var camper = "David";
console.log(camper);
```

上のコードでは、`camper` 変数が最初は `James` として宣言され、その後 `David` に上書きされています。コンソールには文字列 `David` が表示されます。

小さなアプリケーションでは、この種の問題に遭遇しないかもしれません。しかしコードベースが大きくなるにつれて、意図しない変数を誤って上書きしてしまう可能性があります。この動作はエラーをスローしないため、バグの検索や修正が難しくなります。

ES6というJavaScriptの大きなアップデートで、`let` キーワードのこの潜在的な問題を解決するために `var` というキーワードが導入されました。後のチャレンジで他のES6の機能についても学びます。

上のコードで `var` を `let` に置き換えると、エラーになります。

```js
let camper = "James";
let camper = "David";
```

エラーはブラウザのコンソールで確認できます。

つまり、`var` とは異なり、`let` を使うと同じ名前の変数は一度しか宣言できません。

# --instructions--

コードを更新して `let` キーワードのみを使うようにしてください。

# --hints--

コード内に `var` は存在してはいけません。

```js
assert.notMatch(code, /var/g);
```

`catName` は文字列 `Oliver` であるべきです。

```js
assert.equal(catName, 'Oliver');
```

`catSound` は文字列 `Meow!` であるべきです。

```js
assert.equal(catSound, 'Meow!');
```

# --seed--

## --seed-contents--

```js
var catName = "Oliver";
var catSound = "Meow!";
```

# --solutions--

```js
let catName = "Oliver";
let catSound = "Meow!";
```
