---
id: 598e8944f009e646fc236146
title: 関数から返される未定義の値の理解
challengeType: 1
forumTopicId: 301177
dashedName: understanding-undefined-value-returned-from-a-function
---

# --description--

関数には `return` 文を含めることもできますが、必ずしも含める必要はありません。関数に `return` 文がない場合、その関数を呼び出すと、関数は内部のコードを処理しますが、返される値は `undefined` になります。

**例**

```js
let sum = 0;

function addSum(num) {
  sum = sum + num;
}

addSum(3);
```

`addSum` は `return` 文を持たない関数です。この関数はグローバル変数 `sum` を変更しますが、関数の返り値は `undefined` です。

# --instructions--

引数なしの関数 `addFive` を作成してください。この関数は `sum` 変数に 5 を加えますが、返される値は `undefined` です。

# --hints--

`addFive` は関数である必要があります。

```js
assert(typeof addFive === 'function');
```

両方の関数が実行された後、`sum` は `8` と等しくなっているはずです。

```js
assert(sum === 8);
```

`addFive` から返される値は `undefined` であるべきです。

```js
assert(addFive() === undefined);
```

`addFive` 関数の内部で、`5` を `sum` 変数に加える必要があります。

```js
assert(
  __helpers.removeWhiteSpace(addFive.toString()).match(/sum=sum\+5|sum\+=5/)
);
```

# --seed--

## --seed-contents--

```js
// Setup
let sum = 0;

function addThree() {
  sum = sum + 3;
}

// Only change code below this line


// Only change code above this line

addThree();
addFive();
```

# --solutions--

```js
let sum = 0;

function addThree() {
  sum = sum + 3;
}

function addFive() {
  sum = sum + 5;
}

addThree();
addFive();
```
