---
id: 599a789b454f2bbd91a3ff4d
title: 異なる値を比較する練習
challengeType: 1
forumTopicId: 301174
dashedName: practice-comparing-different-values
---

# --description--

前の2つのチャレンジでは、等値演算子 (`==`) と厳密等値演算子 (`===`) について学びました。ここで簡単にレビューし、これらの演算子を使った練習をもう少ししましょう。

比較する値の型が同じでない場合、等値演算子は型変換を行い、その後値を評価します。しかし、厳密等値演算子は型変換をせず、データ型と値の両方をそのまま比較します。

**例**

`3 == '3'` は `true` を返します。これは JavaScript が文字列から数値への型変換を行うためです。`3 === '3'` は `false` を返します。これは型が異なり、型変換が行われないためです。

**注意:** JavaScript では、`typeof` 演算子を使って変数や値の型を調べることができます。使い方は以下の通りです。

```js
typeof 3
typeof '3'
```

`typeof 3` は文字列 `number` を返し、`typeof '3'` は文字列 `string` を返します。

# --instructions--

エディター内の `compareEquality` 関数は等値演算子を使って2つの値を比較しています。値が厳密に等しい場合にのみ文字列 `Equal` を返すように関数を修正してください。

# --hints--

`compareEquality(10, "10")` は文字列 `Not Equal` を返すべきです

```js
assert(compareEquality(10, '10') === 'Not Equal');
```

`compareEquality("20", 20)` は文字列 `Not Equal` を返すべきです

```js
assert(compareEquality('20', 20) === 'Not Equal');
```

`===` 演算子を使うべきです

```js
assert(__helpers.removeJSComments(code).match(/===/g));
```

# --seed--

## --seed-contents--

```js
// Setup
function compareEquality(a, b) {
  if (a == b) { // Change this line
    return "Equal";
  }
  return "Not Equal";
}

compareEquality(10, "10");
```

# --solutions--

```js
function compareEquality(a,b) {
  if (a === b) {
    return "Equal";
  }
  return "Not Equal";
}
```
