---
id: 5cc0bd7a49b71cb96132e54c
title: 再帰を使って数値の範囲を作成する
challengeType: 1
forumTopicId: 301180
dashedName: use-recursion-to-create-a-range-of-numbers
---

# --description--

前のチャレンジに続き、再帰関数を作成して問題を解決する別の機会を提供します。

# --instructions--

2つのパラメータを持つ`rangeOfNumbers`という名前の関数が定義されています。この関数は、`startNum`パラメータで表される数から始まり、`endNum`パラメータで表される数で終わる整数の配列を返す必要があります。開始の数は常に終了の数以下です。関数は自身を呼び出す再帰を使い、ループは一切使わないでください。また、`startNum`と`endNum`が同じ場合にも動作する必要があります。

# --hints--

関数は配列を返す必要があります。

```js
assert(Array.isArray(rangeOfNumbers(5, 10)));
```

コードはループ構文（`for`や`while`、または`forEach`、`map`、`filter`、`reduce`のような高階関数）を使わないでください。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

`rangeOfNumbers`は再帰（自身を呼び出す）を使ってこのチャレンジを解決してください。

```js
assert(
  rangeOfNumbers.toString().match(/rangeOfNumbers\s*\(.+\)/)
);
```

`rangeOfNumbers(1, 5)`は`[1, 2, 3, 4, 5]`を返すべきです。

```js
assert.deepStrictEqual(rangeOfNumbers(1, 5), [1, 2, 3, 4, 5]);
```

`rangeOfNumbers(6, 9)`は`[6, 7, 8, 9]`を返すべきです。

```js
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

`rangeOfNumbers(4, 4)`は`[4]`を返すべきです。

```js
assert.deepStrictEqual(rangeOfNumbers(4, 4), [4]);
```

配列をキャッシュするためにグローバル変数を使わないでください。

```js
rangeOfNumbers(1, 3)
assert.deepStrictEqual(rangeOfNumbers(6, 9), [6, 7, 8, 9]);
```

# --seed--

## --seed-contents--

```js
function rangeOfNumbers(startNum, endNum) {
  return [];
};
```

# --solutions--

```js
function rangeOfNumbers(startNum, endNum) {
  if (endNum - startNum === 0) {
    return [startNum];
  } else {
    const numbers = rangeOfNumbers(startNum, endNum - 1);
    numbers.push(endNum);
    return numbers;
  }
}
```
