---
id: 5cfa3679138e7d9595b9d9d4
title: 再帰を使ってループを置き換える
challengeType: 1
videoUrl: >-
  https://www.freecodecamp.org/news/how-recursion-works-explained-with-flowcharts-and-a-video-de61f40cb7f9/
forumTopicId: 301175
dashedName: replace-loops-using-recursion
---

# --description--

再帰とは、関数が自分自身を使って表現できるという概念です。これを理解するために、まず次の課題を考えてみましょう：配列の最初の`n`個の要素を乗算して、その積を作ることです。`for`ループを使うと、次のようにできます。

```js
  function multiply(arr, n) {
    let product = 1;
    for (let i = 0; i < n; i++) {
      product *= arr[i];
    }
    return product;
  }
```

しかし、`multiply(arr, n) == multiply(arr, n - 1) * arr[n - 1]`ことに気づきます。つまり、`multiply`を自分自身を使って書き換え、ループを使う必要がなくなります。

```js
  function multiply(arr, n) {
    if (n <= 0) {
      return 1;
    } else {
      return multiply(arr, n - 1) * arr[n - 1];
    }
  }
```

`multiply`の再帰バージョンはこのように分解されます。<dfn>基本ケース</dfn>では、`n <= 0`のときに1を返します。`n`が大きい場合は、自分自身を呼び出しますが、`n - 1`を使います。その関数呼び出しは同じ方法で評価され、`multiply`を再び呼び出し、`n <= 0`になるまで続きます。この時点で、すべての関数が戻り、元の`multiply`が答えを返します。

**注意：** 再帰関数は、関数を再度呼び出さずに戻る基本ケースが必要です（この例では`n <= 0`のとき）。そうでないと、実行が終わりません。

# --instructions--

配列`sum(arr, n)`の最初の`n`個の要素の合計を返す再帰関数`arr`を書いてください。

# --hints--

`sum([1], 0)`は0と等しくなるべきです。

```js
assert.equal(sum([1], 0), 0);
```

`sum([2, 3, 4], 1)`は2と等しくなるべきです。

```js
assert.equal(sum([2, 3, 4], 1), 2);
```

`sum([2, 3, 4, 5], 3)`は9と等しくなるべきです。

```js
assert.equal(sum([2, 3, 4, 5], 3), 9);
```

コードはどんな種類のループ（`for`や`while`）や高階関数（`forEach`、`map`、`filter`、`reduce`など）に依存してはいけません。

```js
assert(
  !__helpers.removeJSComments(code).match(/for|while|forEach|map|filter|reduce/g)
);
```

この問題を解くには再帰を使うべきです。

```js
assert(
  sum.toString().match(/sum\(.*\)/g).length > 1
);
```

# --seed--

## --seed-contents--

```js
function sum(arr, n) {
  // Only change code below this line

  // Only change code above this line
}
```

# --solutions--

```js
function sum(arr, n) {
  // Only change code below this line
  if(n <= 0) {
    return 0;
  } else {
    return sum(arr, n - 1) + arr[n - 1];
  }
  // Only change code above this line
}
```
