---
id: 5ee127a03c3b35dd45426493
title: ある変数の値を別の変数に代入する
challengeType: 1
forumTopicId: 418265
dashedName: assigning-the-value-of-one-variable-to-another
---

# --description--

<dfn>代入</dfn>演算子を使って変数に値を代入した後、その変数の値を別の変数に<dfn>代入</dfn>演算子で代入できます。

```js
var myVar;
myVar = 5;
var myNum;
myNum = myVar;
```

上記では、値がない`myVar`変数を宣言し、次に`5`の値を代入しています。次に、`myNum`という名前の変数が値なしで宣言されます。そして、`myVar`（これは`5`）の内容が変数`myNum`に代入されます。これで、`myNum`も`5`の値を持つことになります。

# --instructions--

`a`の内容を変数`b`に代入してください。

# --hints--

指定されたコメントの上のコードは変更しないでください。

```js
assert(/var a;/.test(__helpers.removeJSComments(code)) && /a = 7;/.test(__helpers.removeJSComments(code)) && /var b;/.test(__helpers.removeJSComments(code)));
```

`b`は`7`の値を持つべきです。

```js
assert(typeof b === 'number' && b === 7);
```

`a`は`b`を使って`=`に代入されるべきです。

```js
assert(/b\s*=\s*a\s*/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
// Setup
var a;
a = 7;
var b;

// Only change code below this line
```

# --solutions--

```js
var a;
a = 7;
var b;
b = a;
```
