---
id: cf1111c1c11feddfaeb3bdef
title: JavaScriptで2つの数字を足す
challengeType: 1
forumTopicId: 16650
dashedName: add-two-numbers-with-javascript
---

# --description--

`Number` はJavaScriptのデータ型で、数値データを表します。

では、JavaScriptを使って2つの数字を足してみましょう。

JavaScriptでは、2つの数字の間に`+`記号を置くと、足し算の演算子として使われます。

**例：**

```js
const myVar = 5 + 10;
```

`myVar` は今、`15` の値を持っています。

# --instructions--

`0` を変更して、sumが`20`と等しくなるようにしてください。

# --hints--

`sum` は`20`と等しくなるはずです。

```js
assert(sum === 20);
```

`+` 演算子を使うべきです。

```js
assert(/\+/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
const sum = 10 + 0;
```

# --solutions--

```js
const sum = 10 + 10;
```
