---
id: 69b279bae0f20eda63163cf6
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

У попередньому кроці ви призначили елемент у `leftPart` списку `array`, бо він був меншим. Але так буде не завжди. У деяких випадках порівняння елемент справа може бути меншим.

Створіть умову `else`, яка виконається, коли елемент у `leftPart` не менший за елемент у `rightPart`.

Всередині блоку `else` призначте `rightPart[rightArrayIndex]` `array[sortedIndex]`.

# --hints--

Ви маєте створити блок `else`, який містить `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
