---
id: 69b279bae0f20eda63163cfc
title: Крок 18
challengeType: 1
dashedName: step-18
---

# --description--

Останнє, що потрібно зробити для циклу `while`, — перейти до наступного індексу в відсортованому масиві.

Додайте `1` до значення `sortedIndex`.

Зверніть увагу, що збільшення має відбуватися після призначення елемента лівої частини відсортованому масиву.

# --hints--

Ви повинні збільшити поточне значення `sortedIndex` на одиницю всередині циклу `while`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
