---
id: 69b279bae0f20eda63163cf4
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Коли умова `if` оцінюється як `true`, це означає, що елемент у списку `leftPart` менший за елемент, з яким його порівнюють у списку `rightPart`.

У такому випадку ви можете помістити елемент, знайдений у `leftArrayIndex`, у `leftPart` всередині відсортованого масиву.

У блоці `if` призначте `leftPart[leftArrayIndex]` змінній `array[sortedIndex]`.

# --hints--

Ви повинні призначити `leftPart[leftArrayIndex]` змінній `array[sortedIndex]` всередині `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
