---
id: 69b278519ef46b2fd8bf1546
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У вас уже є значення `middlePoint`. Використайте його, щоб розбити масив на дві половини та підготуйте ліву частину для сортування.

Використайте синтаксис slice, щоб витягти ліву половину `array` і призначити її змінній з ім’ям `leftPart`.

# --hints--

Ви повинні оголосити змінну `leftPart` всередині функції `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Ви повинні призначити `array.slice(0, middlePoint)` змінній `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
