---
id: 69b279bae0f20eda63163cfe
title: Крок 20
challengeType: 1
dashedName: step-20
---

# --description--

Перед тим як тестувати функцію `mergeSort()`, потрібно створити базовий випадок, який зупиняє виконання функції, коли довжина `array` менша або дорівнює `1`.

Цей базовий випадок зупинить рекурсивний виклик. Без нього операція сортування злиттям продовжувала б виконуватися навіть тоді, коли список уже відсортований або не містить елементів.

Одразу після оголошення функції створіть інструкцію `if`, яка перевіряє, чи `array.length` менша або дорівнює 1. У тілі інструкції `if` додайте інструкцію `return`, щоб зупинити виконання функції.

# --hints--

Ви повинні створити інструкцію `if` з умовою `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

У тілі інструкції `return` має бути інструкція `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
