---
id: 69b2762b17d23fbee40a6c00
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Алгоритм сортування злиттям виконує три дії:

- Розбиває несортовану послідовність елементів на підчастини
- Сортує елементи в підчастинах
- Об’єднує відсортовані підчастини

Це відбувається рекурсивно, доки підчастини не об’єднаються у повну відсортовану послідовність. Почнемо з розбиття послідовності.

Спершу, всередині функції `mergeSort`, оголосіть змінну `middlePoint` за допомогою `const` і призначте їй `Math.floor(array.length / 2)`, щоб розбити послідовність навпіл перед рекурсією.

# --hints--

Ви повинні оголосити змінну `middlePoint` і призначити їй `Math.floor(array.length / 2)`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
