---
id: 69b279bae0f20eda63163cf9
title: Крок 15
challengeType: 1
dashedName: step-15
---

# --description--

Цикл `while`, який ви створили, порівнює один елемент з `leftPart` з іншим у `rightPart`, а потім додає менший елемент до головного списку `array`.

Він продовжуватиме цю операцію, доки не залишиться елементів для порівняння. Але в `leftPart` можуть залишатися елементи, тоді як у `rightPart` їх немає, і навпаки.

Створіть ще один цикл `while`, який виконується, коли `leftArrayIndex` менше за довжину `leftPart`. У наступних кроках ви використаєте його, щоб скопіювати залишкові елементи з `leftPart` у список `array`.

# --hints--

Ви повинні створити цикл `while` з умовою `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
