---
id: 69b279bae0f20eda63163cf0
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Внизу тіла вашої функції знову викличте функцію `mergeSort()`. Цього разу передайте `rightPart` як аргумент виклику функції.

# --hints--

Ви маєте викликати `mergeSort()` з аргументом `rightPart` внизу тіла вашої функції.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
