---
id: 69b279bae0f20eda63163cf8
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Інструкції `if` та `else`, які ви створили на попередніх кроках, призначать елементи відсортованому масиву.

Кожен елемент, призначений відсортованому масиву, займає індекс у списку. Тому після кожного присвоєння потрібно перейти до наступного індексу у відсортованому масиві.

Під блоком `if`/`else`, але все ще в циклі `while`, збільшіть `sortedIndex` на `1`.

# --hints--

Ви повинні збільшити поточне значення `sortedIndex` на `1` після вашої умови `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
