---
id: 69b279bae0f20eda63163cef
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер, коли ви розділили список `array` на два окремі списки, продовжуйте ділити кожен список, доки кожен елемент не опиниться у власному списку. Список з одним числом завжди відсортований.

Для цього рекурсивно викликайте `mergeSort` у своїй функції та передавайте `leftPart` як аргумент виклику.

# --hints--

Ви повинні викликати `mergeSort(leftPart)` всередині тієї ж функції `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
