---
id: 69b2713ff6c21fed2d3ad88b
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

Сортування злиттям — це алгоритм сортування, який використовує принцип розділяй і володарюй для сортування колекцій даних. Тобто він «ділить» колекцію на менші підчастини, а потім «перемагає» ці підчастини, сортувавши їх окремо, після чого об’єднує відсортовані підчастини.

Почніть із створення функції з назвою `mergeSort`, яка приймає параметр `array`.

# --hints--

Ви повинні оголосити функцію з ім’ям `mergeSort` з одним параметром `array`.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
