---
id: 69b279bae0f20eda63163cf1
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер настав час відсортувати та об’єднати списки (`leftPart` і `rightPart`) у початковий `array`.

Ви можете зробити це, порівнюючи елементи обох списків і додаючи менший елемент до головного списку. Ви будете робити це порівняння для всіх індексів у `leftPart` і `rightPart`.

Створіть три змінні: `leftArrayIndex`, `rightArrayIndex` і `sortedIndex`, і встановіть їхні значення в `0`. Ці змінні допоможуть вам відстежувати кожен індекс під час процесу сортування.

# --hints--

Ви повинні оголосити змінну `leftArrayIndex`, ініціалізовану `0`.

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

Ви повинні оголосити змінну `rightArrayIndex`, ініціалізовану `0`.

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

Ви повинні оголосити змінну `sortedIndex`, ініціалізовану `0`.

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
