---
id: 69b279bae0f20eda63163cf1
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер настав час відсортувати та об’єднати списки (`leftPart` і `rightPart`) у початковий `array`.

Ви можете зробити це, порівнюючи елементи обох списків і додаючи менший елемент до головного списку. Ви будете робити це порівняння для всіх індексів у `leftPart` і `rightPart`.

Створіть три змінні: `leftArrayIndex`, `rightArrayIndex` і `sortedIndex`, і встановіть їхні значення в `0`. Ці змінні допоможуть вам відстежувати кожен індекс під час процесу сортування.

# --hints--

Ви повинні оголосити змінну `leftArrayIndex`, ініціалізовану `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Ви повинні оголосити змінну `rightArrayIndex`, ініціалізовану `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Ви повинні оголосити змінну `sortedIndex`, ініціалізовану `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
