---
id: 69b278d82f0e514f32acb724
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Після того, як ви вирізали ліву половину, вам все ще потрібна права сторона, щоб завершити крок об’єднання.

Використайте синтаксис slice, щоб витягти праву половину `array` і призначити її змінній з ім’ям `rightPart`.

# --hints--

У вас має бути змінна `rightPart` всередині функції `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Використайте синтаксис slice, щоб витягти праву половину масиву і призначити її змінній `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
