---
id: 69b279bae0f20eda63163cf2
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Далі потрібно створити цикл, який триватиме, доки залишаються елементи в обох `leftPart` і `rightPart`.

Для цього створіть `while` цикл із двома умовами: одна перевіряє, чи `leftArrayIndex` менше за довжину `leftPart`, а інша — чи `rightArrayIndex` менше за довжину `rightPart`.

# --hints--

Ви маєте створити `while` цикл, який перевіряє, чи `leftArrayIndex` менше за довжину `leftPart` і чи `rightArrayIndex` менше за довжину `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
