---
id: 69b279bae0f20eda63163cf3
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

У циклі почніть порівнювати поточні елементи з кожної половини, щоб спочатку об’єднати менше значення.

Додайте блок `if`, який перевіряє, чи поточний елемент у `leftPart` менший за поточний елемент у `rightPart`. 
# --hints--

Умова `if` має перевіряти, чи `leftPart[leftArrayIndex]` менше за `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
