---
id: 69b2762b17d23fbee40a6c00
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Алгоритм сортування злиттям виконує три дії:

- Розділити несортовану послідовність елементів на підчастини
- Відсортувати елементи в підчастинах
- Об’єднати відсортовані підчастини

Вищезазначене відбувається рекурсивно, доки підчастини не об’єднаються у повну відсортовану послідовність. Почнемо з розділення послідовності.

Спершу, всередині функції `mergeSort`, оголосіть змінну `middlePoint` за допомогою `const` і призначте їй `Math.floor(array.length / 2)`, щоб розбити послідовність навпіл перед рекурсією.

# --hints--

Ви повинні оголосити змінну `middlePoint` і призначити їй `Math.floor(array.length / 2)`.

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
