---
id: 69b2713ff6c21fed2d3ad88b
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

Сортування злиттям — це алгоритм сортування, який використовує принцип розділяй і володарюй для сортування колекцій даних. Тобто він "ділить" колекцію на менші підчастини, а потім "володарює" над ними, сортує їх окремо, після чого об’єднує відсортовані підчастини.

Почніть із створення функції з назвою `mergeSort`, яка приймає параметр `array`.

# --hints--

Ви повинні оголосити функцію з ім’ям `mergeSort` з одним параметром `array`.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
