---
id: 68da6138d70960c197ec5ef0
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

이제 메인 BFS 루프를 구현하세요. `while` 루프를 만들어 `queue`이 비어 있지 않은 동안 계속 실행되도록 하세요 (`queue`가 `True`로 평가될 때).

루프 안에서 `queue`를 출력하세요.

# --hints--

`while` 함수 안에 `gen_parentheses` 루프가 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0]
`) })
```

`while` 루프의 조건은 `queue`이어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_conditions()[0].is_equivalent("queue")
`) })
```

`queue` 루프 안에서 `while`를 출력해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_call("print(queue)")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    queue = [('', 0, 0)]
    result = []
    
    
    return result
--fcc-editable-region--
```
