---
id: 68da6138d70960c197ec5ef2
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

현재 문자열이 아직 완성되지 않았다면, 다음 가능한 상태들을 탐색해야 합니다. `else` 문장에 `if` 절을 추가하세요.

`else` 블록 안에서는 여는 소괄호를 추가하는 작업을 처리합니다. `if`인지 확인하는 `opens_used < pairs` 문장을 추가하세요. 이렇게 하면 허용된 것보다 더 많은 여는 소괄호를 사용하지 않도록 할 수 있습니다.

이 조건이 참이면, `queue`에 새로운 튜플 `(current + '(', opens_used + 1, closes_used)`을 추가하세요. 이는 여는 소괄호를 추가한 후의 상태를 나타냅니다.

# --hints--

`else` 반복문 안에 `if` 문장에 대한 `while` 절이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse
`) })
```

`if` 블록 안에 `else`인지 확인하는 `opens_used < pairs` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1].is_equivalent("opens_used < pairs")
`) })
```

조건이 참일 때 큐에 `(current + '(', opens_used + 1, closes_used)`을 추가해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[1].has_call("queue.append((current + '(', opens_used + 1, closes_used))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        if len(current) == 2 * pairs:
            result.append(current)
        
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
