---
id: 68da6138d70960c197ec5eec
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

핵심 알고리즘을 구현하기 전에 입력을 검증해야 합니다. `pairs` 매개변수는 생성할 소괄호 쌍의 개수를 나타내므로 정수여야 합니다.

함수 시작 부분에 `if` 문을 추가하여 `pairs`가 정수가 아닌지 확인하세요. 이를 위해 `isinstance()` 함수를 사용하세요.

조건이 참이면 문자열 `The number of pairs should be an integer`를 반환하세요.

# --hints--

`if` 함수에 `gen_parentheses` 문이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0]
`) })
```

`if` 문은 `pairs`를 사용하여 `not isinstance(pairs, int)`가 정수가 아닌지 확인해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[0].find_conditions()[0].is_equivalent("not isinstance(pairs, int)")
`) })
```

입력이 정수가 아닐 때 문자열 `The number of pairs should be an integer`를 반환해야 합니다.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be an integer'
args = [1.1, "1", [1, 2], {}]
assert  all(gen_parentheses(arg) == msg for arg in args)
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    
    return []
--fcc-editable-region--
```
