---
id: 68da6138d70960c197ec5eed
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

다음으로, 쌍의 개수가 최소한 하나인지 확인해야 합니다. 쌍이 0개이거나 음수이면 소괄호 조합을 생성할 수 없기 때문입니다.

`if` 문장을 하나 더 추가하여 `pairs`이 `1`보다 작은지 확인하세요. 이 조건이 참이면 문자열 `The number of pairs should be at least 1`을 반환하세요.

# --hints--

`if` 함수 안에 두 번째 `gen_parentheses` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_ifs()[1]
`) })
```

두 번째 `if` 문장은 `pairs`이 `1`보다 작은지 확인해야 합니다.

```js
({ test: () => runPython(`
cond = _Node(_code).find_function("gen_parentheses").find_ifs()[1].find_conditions()[0]
assert cond.is_equivalent("pairs < 1") or cond.is_equivalent("1 > pairs")
`) })
```

`The number of pairs should be at least 1`이 `pairs`보다 작을 때 문자열 `1`을 반환해야 합니다.

```js
({ test: () => runPython(`
msg = 'The number of pairs should be at least 1'
assert gen_parentheses(0) == msg and gen_parentheses(-1) == msg
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    
    return []
--fcc-editable-region--
```
