---
id: 68da6138d70960c197ec5ef3
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 닫는 소괄호를 추가하는 작업을 처리해야 합니다. 핵심 규칙은 균형을 유지하는 경우에만 닫는 소괄호를 추가할 수 있다는 점입니다. 즉, 닫는 소괄호보다 여는 소괄호가 더 많이 사용되어야 합니다.

여전히 `else` 블록 안에서, 두 번째 `if` 문장을 추가하여 `closes_used < opens_used`인지 확인하세요.

이 조건이 참이면, 닫는 소괄호를 추가한 후 상태를 나타내는 `queue`라는 새로운 튜플을 `(current + ')', opens_used, closes_used + 1)`에 추가하세요.

# --hints--

`if` 블록 안에 두 번째 `else` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
import ast
assert isinstance(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1], ast.If)
`) })
```

두 번째 중첩된 `if` 문장은 `closes_used < opens_used`인지 확인해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1].test).is_equivalent("closes_used < opens_used")
`) })
```

조건이 참일 때 `(current + ')', opens_used, closes_used + 1)`를 큐에 추가해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].tree.orelse[1]).has_call("queue.append((current + ')', opens_used, closes_used + 1))")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
    
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
        if len(current) == 2 * pairs:
            result.append(current)
--fcc-editable-region--
        else:
            if opens_used < pairs:
                queue.append((current + '(', opens_used + 1, closes_used))
            
--fcc-editable-region--
    return result

print(gen_parentheses(1))
```
