---
id: 68da6138d70960c197ec5eee
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 결과를 저장할 자료 구조를 설정하세요. `result`라는 변수를 선언하고 빈 리스트로 초기화하세요. 이 리스트는 생성하는 모든 유효한 소괄호 조합을 저장할 것입니다.

return 문을 빈 리스트 대신 `result`을 반환하도록 수정하세요.

# --hints--

`result` 함수에서 `gen_parentheses`라는 변수를 선언해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("result")
`) })
```

`result`를 빈 리스트로 초기화해야 합니다.

```js
({ test: () => runPython(`
_result = _Node(_code).find_function("gen_parentheses").find_variable("result")
assert _result.is_equivalent("result = []") or _result.is_equivalent("result = list()")
`) })
```

함수 끝에서 `result`를 반환해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_return("result")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    return []
--fcc-editable-region--
```
