---
id: 68da6138d70960c197ec5eef
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

너비 우선 탐색 방식을 위해, 소괄호 조합을 만들면서 다양한 상태를 추적할 큐를 사용하세요. 각 상태는 세 가지 요소를 포함하는 튜플로 나타냅니다:

- 현재 만들어지고 있는 문자열
- 지금까지 사용한 여는 소괄호 개수
- 지금까지 사용한 닫는 소괄호 개수

`queue`라는 변수를 선언하고, 빈 문자열과 소괄호가 0개인 시작 상태를 나타내는 튜플 `('', 0, 0)`를 포함하는 리스트로 초기화하세요.

# --hints--

`queue` 함수 안에 `gen_parentheses`라는 변수를 선언해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_variable("queue")
`) })
```

`queue`를 튜플 `('', 0, 0)`를 포함하는 리스트로 초기화해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_variable("queue").is_equivalent("queue = [('', 0, 0)]")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    result = []
    return result
--fcc-editable-region--
```
