---
id: 68da604b8b6012a33e2344fb
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 너비 우선 탐색(BFS) 방식을 사용해 모든 유효한 소괄호 조합을 생성하는 함수를 구현할 거예요. 예를 들어, 두 쌍의 소괄호로 만들 수 있는 유효한 조합은 `(())`와 `()()`입니다.

먼저 `gen_parentheses`라는 단일 매개변수를 가진 `pairs`라는 이름의 함수를 만드세요. 지금은 함수에서 빈 리스트를 반환하세요.

# --hints--

`gen_parentheses`라는 이름의 함수를 정의해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_function("gen_parentheses")
`) })
```

`gen_parentheses` 함수는 `pairs`라는 단일 매개변수를 가져야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").has_args("pairs")
`) })
```

함수는 빈 리스트를 반환해야 합니다.

```js
({ test: () => runPython(`
foo = _Node(_code).find_function("gen_parentheses")
assert foo.has_return("[]") or foo.has_return("list()")
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
