---
id: 68da6138d70960c197ec5ef1
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

이제 완전한 소괄호 조합을 만들었는지 확인해야 합니다. 완전한 조합은 쌍의 수의 두 배에 해당하는 길이를 가집니다(각 쌍이 여는 소괄호 하나와 닫는 소괄호 하나를 포함하므로).

`while` 반복문 안에 `if` 문장을 추가하여 `len(current) == 2 * pairs`인지 확인하세요. 이 조건이 참이면 `current`을 `result` 리스트에 추가하세요.

# --hints--

`if` 반복문 안에 `while` 문장이 있어야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0]
`) })
```

`if` 문장은 `len(current) == 2 * pairs`인지 확인해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("len(current) == 2 * pairs")
`) })
```

새로운 `current` 문장 안에서 `result`을 `if` 리스트에 추가해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_call("result.append(current)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)
        
--fcc-editable-region--    
    return result

print(gen_parentheses(1))
```
