---
id: 68e00c5349cbddfef6dbc053
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`while` 반복문 안에서 `queue.pop(0)`을 사용해 큐에서 첫 번째 요소를 제거하고 가져오세요. 이렇게 하면 BFS의 특징인 선입선출(FIFO) 동작을 구현할 수 있습니다.

이 튜플을 `current`, `opens_used`, `closes_used` 세 변수로 풀어내세요.

# --hints--

큐에서 첫 번째 요소를 가져오려면 `queue.pop(0)`를 사용하세요.

```js
({ test: () => runPython(`
assert "queue.pop(0)" in str(_Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0])
`) })
```

결과를 세 변수 `current`, `opens_used`, `closes_used`로 풀어내야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).find_function("gen_parentheses").find_whiles()[0].find_bodies()[0].has_stmt("current, opens_used, closes_used = queue.pop(0)")
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []
--fcc-editable-region--
    while queue:
        print(queue)
        
--fcc-editable-region--
    return result
```
