---
id: 68e00f88cf82083df8e0497b
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

BFS 로직 작업을 계속하기 전에 `gen_parentheses(1)`를 호출하고 결과를 콘솔에 출력하세요.

# --hints--

`gen_parentheses(1)`를 콘솔에 출력해야 합니다.

```js
({ test: () => runPython(`
assert _Node(_code).has_call('print(gen_parentheses(1))')
`) })
```

# --seed--

## --seed-contents--

```py
def gen_parentheses(pairs):
    if not isinstance(pairs, int):
        return 'The number of pairs should be an integer'
    if pairs < 1:
        return 'The number of pairs should be at least 1'
    
    queue = [('', 0, 0)]
    result = []

    while queue:
        print(queue)
        current, opens_used, closes_used = queue.pop(0)

    return result
--fcc-editable-region--

--fcc-editable-region--
```
