---
id: 66f686b8ebdb982fa8e14330
title: crea un programma di conversione delle date
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

In questo laboratorio, scriverai un programma di conversione delle date che trasforma una data data in formati diversi. Per esempio, la data attuale `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` sarebbe formattata nei seguenti 2 modi:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti creare una variabile chiamata `currentDate` e assegnarle la data e l'ora correnti usando l'oggetto `Date`.
2. Dovresti creare una variabile chiamata `currentDateFormat` e assegnarle la stringa `Current Date and Time: [current date]`. Sostituisci `[current date]` con il risultato della variabile `currentDate`.
3. Dovresti stampare il valore di `currentDateFormat` nella console.
4. Dovresti creare una funzione chiamata `formatDateMMDDYYYY` che prende l'oggetto data come parametro. Puoi chiamare questo parametro come preferisci.
5. La tua funzione `formatDateMMDDYYYY` dovrebbe convertire la data corrente nel formato `month/day/year` e restituire la stringa `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Dovresti creare una funzione chiamata `formatDateLong` che prende l'oggetto data come parametro.
7. La tua funzione `formatDateLong` dovrebbe convertire la data corrente nel formato `Month Day, Year` e restituire la stringa `Formatted Date (Month Day, Year): [formatted date]`.

**Nota:** Per far passare i test, assicurati di usare `en-US` come locale quando formatti le date.

# --hints--

Dovresti avere una variabile chiamata `currentDate` che contiene la data e l'ora correnti usando l'oggetto `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Dovresti avere una variabile chiamata `currentDateFormat` che contiene la data corrente nel formato `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Dovresti stampare il valore di `currentDateFormat` nella console.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Dovresti avere una funzione `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

La funzione `formatDateMMDDYYYY` dovrebbe prendere un singolo parametro.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Quando l'oggetto data contiene `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la funzione `formatDateMMDDYYYY` dovrebbe restituire `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Dovresti avere una funzione `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

La funzione `formatDateLong` dovrebbe prendere un singolo parametro.

```js
assert.lengthOf(formatDateLong, 1);
```

Quando l'oggetto data contiene `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la funzione `formatDateLong` dovrebbe restituire `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
