---
id: 69a92f0b9cc04bb0d5327bb5
title: Einen Stack implementieren
challengeType: 26
dashedName: implement-a-stack
---

# --description--

In diesem Labor implementieren Sie eine Stack-Datenstruktur mithilfe von Funktionen. Ein Stack ist eine Last-In-First-Out-(LIFO-)Datenstruktur, bei der Elemente oben hinzugefügt und entfernt werden.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine `initStack`-Funktion haben, die ein Objekt mit einer `collection`-Eigenschaft zurückgibt, die auf ein leeres Array gesetzt ist.
1. Sie sollten eine `push`-Funktion haben, die ein Element oben auf den Stack legt.
1. Sie sollten eine `pop`-Funktion haben, die das oberste Element des Stacks entfernt und zurückgibt oder `undefined`, falls keines vorhanden ist.
1. Sie sollten eine `peek`-Funktion haben, die das oberste Element des Stacks zurückgibt, ohne es zu entfernen, oder `undefined`, falls keines vorhanden ist.
1. Sie sollten eine `isEmpty`-Funktion haben, die `true` zurückgibt, wenn der Stack keine Elemente enthält, und `false` sonst.
1. Sie sollten eine `clear`-Funktion haben, die alle Elemente aus dem Stack entfernt.

**Hinweis:** Die meisten Tests hängen von den `initStack`- und `push`-Funktionen ab. Implementieren Sie diese zuerst, da Tests für `pop`, `peek`, `isEmpty` und `clear` das Hinzufügen von Elementen zum Stack erfordern.

# --hints--

Sie sollten eine `push`-Funktion haben.

```js
assert.isFunction(push);
```

Die `push`-Funktion sollte ein Element oben auf den Stack legen.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

Sie sollten eine `pop`-Funktion haben.

```js
assert.isFunction(pop);
```

Die `pop`-Funktion sollte das oberste Element des Stacks entfernen und zurückgeben.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.strictEqual(pop(stack), 'second');
```

Die `pop`-Funktion sollte `undefined` zurückgeben, wenn der Stack leer ist.

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

Die `pop`-Funktion sollte falsy-Werte korrekt zurückgeben.

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

Sie sollten eine `peek`-Funktion haben.

```js
assert.isFunction(peek);
```

Die `peek`-Funktion sollte das oberste Element des Stacks zurückgeben, ohne es zu entfernen.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
```

Die `peek`-Funktion sollte `undefined` zurückgeben, wenn der Stack leer ist.

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

Die `peek`-Funktion sollte falsy-Werte korrekt zurückgeben.

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

Sie sollten eine `isEmpty`-Funktion haben.

```js
assert.isFunction(isEmpty);
```

Die `isEmpty`-Funktion sollte `true` für einen leeren Stack zurückgeben.

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

Die `isEmpty`-Funktion sollte `false` für einen nicht-leeren Stack zurückgeben.

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

Die `isEmpty`-Funktion sollte `false` zurückgeben, wenn das oberste Element falsy ist.

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

Sie sollten eine `clear`-Funktion haben.

```js
assert.isFunction(clear);
```

Die `clear`-Funktion sollte alle Elemente aus dem Stack entfernen.

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
