---
id: 61537c9eecea6a335db6da79
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Chini ya kipengele chako cha `.header`, tengeneza kipengele kipya cha `div` na weka thamani ya `class` kuwa `gallery`. `div` hii itatumika kama kontena la picha za maktaba.

Ndani ya kipengele hicho cha `.gallery`, tengeneza vipengele tisa vya `img`.

# --hints--

Unapaswa kutengeneza kipengele cha `div` ndani ya kipengele chako cha `body`.

```js
assert(document.querySelector('body')?.querySelectorAll('div')?.length >= 1);
```

Kipengele chako kipya cha `div` kinapaswa kuwa na `class` yenye thamani ya `gallery`.

```js
assert(document.querySelector('body')?.querySelector('.gallery'));
```

Kipengele chako kipya cha `div` kinapaswa kuja baada ya kipengele chako cha `.header`.

```js
assert(document.querySelector('header')?.nextElementSibling?.classList?.contains('gallery'));
```

Kipengele chako cha `.gallery` kinapaswa kuwa na vipengele tisa vya `img`.

```js
assert(document.querySelector('.gallery')?.querySelectorAll('img')?.length === 9);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
  </body>
</html>
--fcc-editable-region--
```

```css

```
