---
id: 61537485c4f2a624f18d7794
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Anza na muundo wako wa msingi wa HTML. Ongeza tamko la `DOCTYPE`, kipengele cha `html`, kipengele cha `head`, na kipengele cha `body`.

Ongeza sifa ya kipengele ya `lang` kwenye lebo ya kufungua ya `<html>` ukiwa na `en` kama thamani.

# --hints--

Msimbo wako unapaswa kuwa na rejea ya `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kuweka nafasi baada ya rejea ya `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua `html` kama aina ya hati.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa lebo ya kufunga ya `>` baada ya aina.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufungua.

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Leboa ya kufungua ya `<html>` inapaswa kuwa na sifa ya kipengele ya `lang` yenye thamani ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `head`.

```js
assert(code.match(/<head\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `head`.

```js
assert(code.match(/<\/head\s*>/i));
```

Unapaswa kuwa na lebo ya kufungua ya `body`.

```js
assert(code.match(/<body\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `body`.

```js
assert(code.match(/<\/body\s*>/i));
```

Vipengele vya `head` na `body` vinapaswa kuwa ndugu.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

Kipengele cha `head` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Kipengele cha `body` kinapaswa kuwa ndani ya kipengele cha `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
