---
id: 64dd076a1b33c86d84b4232f
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Ili kuona vizuri jinsi vipengele vyako vinavyopimwa, kuongeza mpaka kunaweza kusaidia.

Toa `.gallery` kipengele chako `width` ya `50%` na `border` iliyowekwa kuwa `5px solid red`.

Kisha toa vipengele vyako vya `img` `width` ya `100%`, `padding` iliyowekwa kuwa `5px`, na `border` iliyowekwa kuwa `5px solid blue`.

# --hints--

Unapaswa kuwa na kichaguzi cha `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Kichaguzi chako cha `.gallery` kinapaswa kuwa na sifa ya `border` iliyowekwa kuwa `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Unapaswa kuwa na kichaguzi cha `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `width` iliyowekwa kuwa `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `padding` iliyowekwa kuwa `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Kichaguzi chako cha `img` kinapaswa kuwa na sifa ya `border` iliyowekwa kuwa `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
