---
id: 65df3afd233057f6a620a860
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Dies zeigt, wie die Hilfefunktionen verwendet werden können, um die Klassendefinitionen eines Lernenden zu analysieren

# --instructions--

Wir können die Klassendefinitionen in Python finden, gebe es an das JavaScript zurück, und prüfe dann, ob sie die Zeichenfolge „pass“ enthält.

Wenn wir etwas zuverlässigeres benötigen, können wir die Hilfsprogramme erweitern, um spezifisch nach Pass-Anweisungen zu suchen. Dies wäre nützlich, wenn wir z. B. Klassen schreiben wollen, die z. B. Passagiere enthalten.

# --hints--

`TreeNode` sollte keine `pass` Anweisungen enthalten.

```js
({ test: () => {
  const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
  assert.notInclude(pyClassStr, "pass")
}})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass

class Trap:
    pass
--fcc-editable-region--
```

# --solutions--

```py
class TreeNode:
    def __init__(self, thing):
        self.thing = thing

class Trap:
    pass

```
