---
id: 69b83e35f19ba26ba1fa517a
title: 交通信号シーケンサーを作成する
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

このラボでは、設定可能な交通信号のサイクルをシミュレートし、異常を記録します。

交通信号のフェーズを表す設定オブジェクトを扱います。各設定オブジェクトには次のプロパティがあります。

- `fault`: `true` の場合に早期終了を引き起こすブール値のフラグ。
- `phases`: フェーズオブジェクトの配列。

`phases` 内の各フェーズオブジェクトには次のプロパティがあります。

- `color`: 信号の色を表す文字列（`"green"`、`"yellow"`、または `"red"`）。
- `duration`: フェーズが続く秒数を表す正の整数。

提供されている `config1`、`config2`、`config3`、`config4` のオブジェクトは、設定オブジェクトの例として参照できます。

**注意**: 追加の `console.log()` ステートメントを入れないでください。テストが失敗する場合があります。

**目的**：以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

# ユーザーストーリー

1. `runSequence` という名前の関数を用意してください。パラメータは `config` と `cycles` の2つで、`config` は設定オブジェクトを、`cycles` はシーケンスが停止するまでの最大実行回数を表します。

1. `runSequence(config, cycles)` を `for` または `while` ループを使って、指定された回数のサイクルにわたり各フェーズを繰り返し処理するように実装してください。

1. `runSequence` 関数は以下の動作をする必要があります。
    - `config.phases` が空の場合は `No phases found` とログに出力し、すぐに戻る。
    - `config.fault` が `true` の場合は `Faulted phase!` とログに出力し、シミュレーションを早期終了する。
    - `duration <= 0` の場合は `Invalid phase detected` とログに出力する。
    - 有効なフェーズごとに `Switching to [color] for [duration] s` とログに出力する。`[color]` と `[duration]` はフェーズオブジェクトの対応するプロパティに置き換える。
    - 例えば、`runSequence(config1, 1)` は次のようにログに出力します。

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. `generateTimeline` という名前の関数を用意してください。パラメータは `config` と `cycles` の2つです。

1. `generateTimeline` 関数は以下の動作をする必要があります。
    - サイクルを通じて各フェーズの経過時間の累積を配列に記録し、繰り返し処理しながら各フェーズの `duration` を合計に加える。
    - `config.fault` が `true` であっても、また `duration <= 0` であっても、すべての故障フェーズや無効なフェーズをバリデーションなしで処理する。
    - 累積タイムスタンプの配列を返す。
    - 例えば、`generateTimeline(config1, 1)` は `[5, 7, 11]` の配列を返します。

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

`config1`、`config2`、`config3`、`config4` のオブジェクトは変更しないでください。

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

`runSequence` という名前の関数を用意してください。

```js
assert.isFunction(runSequence);
```

`runSequence` 関数は `config` と `cycles` の2つのパラメータを受け取る必要があります。

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` は順に `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s` とログに出力します。

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` は順に `Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s`、`Switching to green for 5 s`、`Switching to yellow for 2 s`、`Switching to red for 4 s` とログに出力します。

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` は順に `Switching to red for 3 s`、`Invalid phase detected`、`Switching to green for 6 s` とログに出力します。

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` は `Faulted phase!` とログに出力します。

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` は `No phases found` とログに出力します。

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

`generateTimeline` という名前の関数を用意してください。

```js
assert.isFunction(generateTimeline);
```

`generateTimeline` 関数は `config` と `cycles` の2つのパラメータを受け取る必要があります。

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` は `[5, 7, 11]` を返します。

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` は `[5, 7, 11, 16, 18, 22]` を返します。

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` は `[3, 1, 7, 10, 8, 14]` を返します。

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` は `[5, 7, 13]` を返します。

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` は `[]` を返します。

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
