---
id: 68a47dcdbce12193caf77012
title: ソート済みインデックスファインダーを実装する
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

このラボでは、値を昇順にソートした配列に挿入すべき最も低いインデックスを返す関数を作成します。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `getIndexToIns` 関数は2つの引数を取る必要があります：配列と数値。  
2. `sort` メソッドを使って配列を昇順にソートしてください。
3. `getIndexToIns` 関数は `findIndex` メソッドを使って数値を挿入すべき最も低いインデックスを返す必要があります。
4. `getIndexToIns` 関数は常に数値を返す必要があります。

**ヒント：**  
`findIndex` メソッドはJavaScriptの組み込み配列メソッドです。コールバック関数を取り、条件を満たす最初の要素のインデックスを返します。`findIndex` と `sort` はどちらも高階関数です。

**例：**

- `getIndexToIns([1, 2, 3, 4], 1.5)` は `1` を返すべきです。なぜなら `1.5` は `1`（インデックス `0`）より大きく、`2`（インデックス `1`）より小さいからです。
- `getIndexToIns([20, 3, 5], 19)` は `2` を返すべきです。なぜなら `[3, 5, 20]` にソートした後、`19` は `20`（インデックス `2`）より小さく、`5`（インデックス `1`）より大きいからです。


# --hints--

`getIndexToIns` 関数を用意してください。

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` は常に数値を返す必要があります。

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

関数は `sort` メソッドを使うべきです。

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


関数は `findIndex` メソッドを利用するべきです。

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` は `3` を返すべきです。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` は `1` を返すべきです。

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` は `0` を返すべきです。

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` は `2` を返すべきです。

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` は `3` を返すべきです。

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` は `0` を返すべきです。

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
