---
id: 587d8259367417b2b2512c86
title: 삽입 정렬 구현하기
challengeType: 26
dashedName: implement-insertion-sort
---

# --description--

이 실습에서는 삽입 정렬 알고리즘을 구현할 것입니다. 이 메서드는 리스트 앞부분에 정렬된 배열을 점차 만들어 나가는 방식으로 작동합니다. 먼저 첫 번째 요소를 정렬된 배열로 시작합니다. 그다음 다음 요소를 검사하여 정렬된 배열 안에서 올바른 위치에 들어갈 때까지 뒤로 교환합니다. 리스트 끝에 도달할 때까지 계속해서 새로운 항목을 정렬된 부분으로 뒤로 교환하며 반복합니다. 이 알고리즘은 평균 및 최악의 경우에 이차 시간 복잡도를 가집니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `insertionSort` 함수를 만들어야 합니다.
1. `insertionSort` 함수는 정수 배열을 입력받아 같은 정수들을 오름차순으로 정렬한 배열을 반환해야 합니다.

# --hints--

`insertionSort`는 함수여야 합니다.

```js
assert.isFunction(insertionSort);
```

`insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])`는 순서만 바뀌었을 뿐 변경되지 않은 배열을 반환해야 합니다.

```js
assert.sameMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`insertionSort`는 오름차순으로 정렬된 배열을 반환해야 합니다.

```js

assert.sameOrderedMembers(
  insertionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
);
```

`insertionSort([5, 4, 33, 2, 8])`는 `[2, 4, 5, 8, 33]`를 반환해야 합니다.

```js
assert.sameOrderedMembers(insertionSort([5, 4, 33, 2, 8]), [2, 4, 5, 8, 33])
```

`insertionSort`는 내장 `.sort()` 메서드를 사용하지 않아야 합니다.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    insertionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function insertionSort (array) {
  for (let currentIndex = 0; currentIndex < array.length; currentIndex++) {
    let current = array[currentIndex];
    let j = currentIndex - 1;
    while (j > -1 && array[j] > current) {
      array[j + 1] = array[j];
      j--;
    }
    array[j + 1] = current;
  }
  return array;
}
```
