---
id: 67ed03ac474c48692f41749e
title: ハッシュテーブルを作成する
challengeType: 27
dashedName: build-a-hash-table
saveSubmissionToDB: true
---

# --description--

このラボでは、ハッシュテーブルを一から作成します。ハッシュテーブルはキーと値のペアをストアするデータ構造です。ハッシュテーブルはキーを入力として受け取り、特定のハッシュ関数に従ってこのキーをハッシュ化することで動作します。

このラボの目的のために、ハッシュ関数は単純で、キーの各文字のUnicode値の合計を計算します。ハッシュ値は、関連する値をストアするための実際のキーとして使われます。同じハッシュ値は、キーに関連付けられた値を取得および削除する際にも使われます。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `HashTable` という名前のクラスを定義し、新しい`collection`のインスタンスが作成されるときに、空の辞書に初期化された`HashTable`属性を持つべきです。`collection`辞書は、キーのハッシュ値に基づいてキーと値のペアをストアします。

2. `HashTable`クラスは、`hash`、`add`、`remove`、`lookup`の4つのインスタンスメソッドを持つべきです。

3. `hash`メソッドは以下を行うべきです：

    - 文字列をパラメータとして受け取る。
    - 文字列の各文字のUnicode（ASCII）値の合計として計算されたハッシュ値を返す。計算には`ord`関数を使えます。

4. `add`メソッドは以下を行うべきです：
    - キーと値のペアを表す2つの引数を受け取り、キーのハッシュを計算する。
    - 計算したハッシュ値をキーとして使い、`collection`辞書の中にキーと値のペアを含む辞書をストアする。
    - 複数のキーが同じハッシュ値を生成する場合、それらのキーと値のペアは同じハッシュ値の下にある既存のネストされた辞書にストアされるべきです。

5. `remove`メソッドは以下を行うべきです：

    - キーを引数として受け取り、そのハッシュを計算する。
    - コレクションにキーが存在するか確認する。
    - ハッシュテーブルから対応するキーと値のペアを削除する。
    - キーがコレクションに存在しない場合は、エラーを発生させたり何も削除したりしない。

6. `lookup`メソッドは以下を行うべきです：

    - キーを引数として受け取る。
    - キーのハッシュを計算し、ハッシュテーブル内にストアされた対応する値を返す。
    - キーがコレクションに存在しない場合は`None`を返す。

# --hints--

`HashTable`クラスを定義してください。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).has_class("HashTable")
    `)) 
})
```

新しい`HashTable`クラスのインスタンスが作成されるとき、その`collection`属性は空の辞書に初期化されるべきです。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.collection == {}
    `) 
})
```

`HashTable`クラスは`hash`メソッドを持つべきです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('hash')`))}})
```

`hash`メソッドは文字列をパラメータとして受け取るべきです。

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.hash)
    len(sig.parameters) == 1
  `))
})
```

`hash`メソッドは文字列を引数として受け取り、その文字列の各文字のUnicode値の合計を返すべきです。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("fcc")
        assert hash_result == 300

        # prevent hardcoding
        assert ht.hash("golf") == 424
        assert ht.hash("read") == 412
`) })
```

`HashTable`クラスは`add`メソッドを持つべきです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('add')`))}})
```

`add`メソッドはキーと値をパラメータとして受け取るべきです。

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.add)
    len(sig.parameters) == 2
  `))
})
```

`HashTable`クラスは`remove`メソッドを持つべきです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('remove')`))}})
```

`remove`メソッドはキーをパラメータとして受け取るべきです。

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.remove)
    len(sig.parameters) == 1
  `))
})
```

コレクションに存在しないキーを削除しようとしたとき、エラーを発生させたり何も削除したりしてはいけません。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("rose", "flower")
        index = ht.hash("rose")
        original = ht.collection.copy()

        ht.remove("tulip")
        ht.remove("sore")

        assert ht.collection == original
        assert "rose" in ht.collection[index]
        assert "tulip" not in ht.collection.get(index, {})
    `) 
})
```

複数のキーが同じインデックスにハッシュされる場合、`remove`メソッドはそのインデックスの辞書全体ではなく、特定のキーと値のペアだけを削除するべきです。

```js
({
  test: () => runPython(`
    ht = HashTable()
    ht.add("rose", "flower")
    ht.add("sore", "pain")  # "rose" and "sore" both hash to the same index

    index = ht.hash("rose")
    ht.remove("rose")

    assert index in ht.collection
    assert "rose" not in ht.collection[index]
    assert "sore" in ht.collection[index]
    assert ht.collection[index]["sore"] == "pain"
  `)
})
```

`HashTable`クラスは`lookup`メソッドを持つべきです。

```js
({test: () => {assert(runPython(`_Node(_code).find_class('HashTable').has_function('lookup')`))}})
```

`lookup`メソッドはキーをパラメータとして受け取るべきです。

```js
({ test: () => assert(runPython(`
    import inspect 
    ht= HashTable()
    sig = inspect.signature(ht.lookup)
    len(sig.parameters) == 1
  `))
})
```

`HashTable().hash('golf')`は`424`を返すべきです。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        hash_result = ht.hash("golf")
        assert hash_result == 424

        # prevent hardcoding
        assert ht.hash("dear") == 412
        assert ht.hash("cat") == 312
    `) 
})
```

`HashTable().add('golf', 'sport')`はキーと値のペアをキー`424`のコレクションに追加するべきです。

```js

({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")
    expected_value = {424: {'golf': 'sport'}}

    assert ht.collection == expected_value
    `) 
})
```

`HashTable().add('dear', 'friend')`と`HashTable().add('read', 'book')`は、キー`412`のインデックスにネストされた辞書として両方のキーと値のペアをコレクションに追加するべきです。

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("dear", "friend")
    ht.add("read", "book")

    expected_value = {
        "dear": "friend",
        "read": "book"
    }

    assert ht.collection.get(412) == expected_value
    `) 
})
```

キーがハッシュテーブルに存在するとき、`remove()`メソッドは指定されたキーと対応する値をコレクションから削除するべきです。

```js
({ 
    test: () => runPython(`
    ht = HashTable()
    ht.add("golf", "sport")

    expected_value_before_removal = {
        "golf": "sport"
    }

    index = ht.hash("golf")

    assert ht.collection.get(index) == expected_value_before_removal

    ht.remove("golf")

    assert "golf" not in ht.collection.get(index, {})
    `) 
})
```

`'golf', 'sport'`キーと値のペアがハッシュテーブルに存在するとき、`HashTable().lookup('golf')`は`sport`を返すべきです。

```js
({ 
    test: () => runPython(`
    ht = HashTable()

    ht.add("golf", "sport")

    expected_value = "sport"
    assert ht.lookup("golf") == expected_value
    `) 
})

```

`'golf', 'sport'`キーと値のペアがコレクションに存在しないとき、`HashTable().lookup('golf')`は`None`を返すべきです。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        assert ht.lookup("golf") is None
`) })

```

`'fcc'`キーがコレクションに存在するとき、`HashTable().lookup('cfc')`は`None`を返すべきです。

```js
({ 
    test: () => runPython(`
        ht = HashTable()
        ht.add("fcc", "coding")
        assert ht.lookup("cfc") is None
`) })
```

`('rose', 'flower')`をハッシュテーブルに追加するとき、その`collection`属性はこのようになります：`{ 441: { 'rose': 'flower' }}`。

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("rose", "flower")

    expected_value = {
        "rose": "flower"
    }

    assert ht.collection.get(441) == expected_value

    # prevent hardcoding

    ht2 = HashTable()

    ht2.add("kebab", "food")
    expected_value = {
        "kebab": "food"
    }

    assert ht2.collection.get(501) == expected_value
    `)
})
```

既存のキーと同じハッシュ値にハッシュされるキーを追加するとき、例えば`fcc`と`cfc`、`collection`はこのようになります：`{ 300: { 'fcc': 'coding', 'cfc':  'chemical' }}`。

```js
({ 
    test: () => runPython(`
 
    ht = HashTable()

    ht.add("fcc", "coding")
    ht.add("cfc", "chemical")

    expected_value = {
        "fcc": "coding",
        "cfc": "chemical"
    }

    assert ht.collection.get(300) == expected_value

    # prevent hardcoding

    ht2 = HashTable()
    ht2.add("cat", "animal")
    ht2.add("act", "verb")

    expected_value = {
        "cat": "animal",
        "act": "verb"
    }

    assert ht2.collection.get(312) == expected_value
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class HashTable:
    def __init__(self):
        self.collection = {}

    def hash(self, string):
        hashed = 0
        for char in string:
            hashed += ord(char)  
        return hashed

    def add(self, key, val):
        the_hash = self.hash(key)
        if the_hash not in self.collection:
            self.collection[the_hash] = {}
        self.collection[the_hash][key] = val

    def remove(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            del self.collection[the_hash][key]
            if not self.collection[the_hash]: 
                del self.collection[the_hash]

    def lookup(self, key):
        the_hash = self.hash(key)
        if the_hash in self.collection and key in self.collection[the_hash]:
            return self.collection[the_hash][key]
        return None
```
