---
id: 5688e62ea601b2482ff8422b
title: بناء بحث في الملف الشخصي
challengeType:  26
dashedName: lab-profile-lookup
---

# --description--

في هذا الدرس، ستبني بحثًا في الملف الشخصي يستخرج معلومات عن الأشخاص في قائمة جهات الاتصال.

لنفترض في هذا المثال وجود جهة اتصال باسم Akira Laine، يجب أن تُرجع الدالة `lookUpProfile("Akira", "lastName")` القيمة `Laine`.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال الدرس.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `lookUpProfile` تأخذ اسمًا وخاصية كوسيطتين.
2. يجب أن تسترجع معلومات جهة الاتصال من مصفوفة `contacts` المتوفرة.
3. إذا تلقّت الدالة اسم جهة اتصال وكانت الخاصية موجودة على جهة الاتصال المرتبطة، فينبغي إعادة قيمة الخاصية.
4. إذا لم يطابق الاسم المُمرَّر إلى الدالة أي جهة اتصال في مصفوفة جهات الاتصال، فيجب أن تُعيد الدالة `"No such contact"`.
5. إذا لم تكن الخاصية موجودة على جهة اتصال تم العثور عليها، فيجب أن تُعيد الدالة `"No such property"`.

# --before-each--

```js
const _contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];

globalThis.contacts = structuredClone(_contacts);
```

# --hints--

يجب أن تبقى مصفوفة `contacts` موجودة في النطاق العام. أعد ضبط الدرس لاستعادتها إذا حذفتها.

```js
assert.match(code, /(?:let|const|var)\s+contacts\s*=\s*\[/);
```

يجب أن يكون لديك دالة باسم `lookUpProfile`.

```js
assert.isFunction(lookUpProfile);
```

يجب أن تُرجع `lookUpProfile("Kristian", "lastName")` سلسلة نصية.

```js
assert.isString(lookUpProfile("Kristian", "lastName"));
```

يجب أن تُرجع `lookUpProfile("Kristian", "lastName")` السلسلة النصية `Vos`

```js
assert.strictEqual(lookUpProfile('Kristian', 'lastName'), 'Vos');
```

يجب أن تُرجع `lookUpProfile("Sherlock", "likes")` المصفوفة `["Intriguing Cases", "Violin"]`

```js
assert.deepEqual(lookUpProfile('Sherlock', 'likes'), [
  'Intriguing Cases',
  'Violin'
]);
```

يجب أن تُرجع `lookUpProfile("Harry", "likes")` مصفوفة

```js
assert.isArray(lookUpProfile('Harry', 'likes'));
```

يجب أن تُرجع `lookUpProfile("Bob", "number")` السلسلة النصية `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'number'),'No such contact');
```

يجب أن تُرجع `lookUpProfile("Bob", "potato")` السلسلة النصية `No such contact`

```js
assert.strictEqual(lookUpProfile('Bob', 'potato'), 'No such contact');
```

يجب أن تُرجع `lookUpProfile("Akira", "address")` السلسلة النصية `No such property`

```js
assert.strictEqual(lookUpProfile('Akira', 'address'), 'No such property');
```

يجب أن تبقى مصفوفة `contacts` دون تغيير بعد تشغيل `lookUpProfile`

```js
lookUpProfile("Sherlock", "likes");
assert.deepEqual(contacts,_contacts);
```

# --seed--

## --seed-contents--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


```

# --solutions--

```js
let contacts = [
  {
    firstName: "Akira",
    lastName: "Laine",
    number: "0543236543",
    likes: ["Pizza", "Coding", "Brownie Points"],
  },
  {
    firstName: "Harry",
    lastName: "Potter",
    number: "0994372684",
    likes: ["Hogwarts", "Magic", "Hagrid"],
  },
  {
    firstName: "Sherlock",
    lastName: "Holmes",
    number: "0487345643",
    likes: ["Intriguing Cases", "Violin"],
  },
  {
    firstName: "Kristian",
    lastName: "Vos",
    number: "unknown",
    likes: ["JavaScript", "Gaming", "Foxes"],
  },
];


function lookUpProfile(name, prop) {
  for (let i in contacts) {
    if (contacts[i].firstName === name) {
      return contacts[i][prop] || "No such property";
    }
  }
  return "No such contact";
}
```
