---
id: 68e4276483b0c71658aab1cc
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

あなたの `userInput` の値は、文字列の両端に余分な空白があります。その空白を取り除いた新しい文字列を作成するには、`trim()` メソッドを使うことができます。

例を示します：

```js
const sentence = "  Learning about the trim method.  ";
sentence.trim(); // "Learning about the trim method."
```

`cleanedInput` という名前の変数を作成し、`trim()` 変数に対して `userInput` メソッドを使った結果を代入してください。

# --hints--

`cleanedInput` という変数があるはずです。

```js
assert.exists(cleanedInput);
```

`cleanedInput` 変数は文字列であるべきです。

```js
assert.isString(cleanedInput);
```

`userInput.trim()` を `cleanedInput` 変数に代入してください。

```js
assert.equal(cleanedInput, userInput.trim());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);
--fcc-editable-region--

--fcc-editable-region--
```
