---
id: 68e4276483b0c71658aab1ce
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

文字列の先頭の空白だけを取り除きたい場合があります。そんなときに `trimStart()` メソッドが役立ちます。

例を示します：

```js
const sentence = "  I love learning.  ";
sentence.trimStart(); // "I love learning.  "
```

文字列は不変であることを覚えておいてください。このメソッドは結果を持つ新しい文字列を作成します。元の文字列は変更しません。

変数 `trimmedStart` を作成し、`trimStart()` 変数に対して `userInput` メソッドを使った結果を代入してください。

# --hints--

`trimmedStart` という変数を持っているはずです。

```js
assert.exists(trimmedStart);
```

`trimmedStart` 変数は文字列であるべきです。

```js
assert.isString(trimmedStart);
```

`userInput.trimStart()` を `trimmedStart` 変数に代入してください。

```js
assert.equal(trimmedStart, userInput.trimStart());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);
--fcc-editable-region--

--fcc-editable-region--
```
