---
id: 68e4276483b0c71658aab1ca
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、文字列のフォーマットに使われるさまざまな文字列メソッドを使う練習をします。

まず、`userInput` という変数を作成し、次の値を代入してください：`"   Hello World!   "`。

# --hints--

`userInput` という変数があるはずです。

```js
assert.exists(userInput);
```

`userInput` 変数は文字列であるべきです。

```js
assert.isString(userInput);
```

`"   Hello World!   "` を `userInput` 変数に代入してください。`H` の前に3つのスペース、`!` の後にも3つのスペースがあることを確認してください。

```js
assert.equal(userInput, "   Hello World!   ");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
