---
id: 68e4276483b0c71658aab1d1
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

すべての文字が小文字の新しい文字列を作成する必要がある場合は、`toLowerCase()` メソッドを使うことができます。

まず、`lowerCaseInput` という変数を作成し、`toLowerCase()` 変数に対して `cleanedInput` メソッドを使った結果を代入してください。

次に、文字列 `"Result of using the toLowerCase() method:"` をコンソールに出力してください。

最後に、`console.log()` 変数をもう一度 `lowerCaseInput` してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`lowerCaseInput` という変数を用意してください。

```js
assert.exists(lowerCaseInput);
```

`lowerCaseInput` 変数は文字列である必要があります。

```js
assert.isString(lowerCaseInput);
```

`cleanedInput.toLowerCase()` を `lowerCaseInput` 変数に代入してください。

```js
assert.equal(lowerCaseInput, cleanedInput.toLowerCase());
```

文字列 `"Result of using the toLowerCase() method:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[10], "Result of using the toLowerCase() method:");
```

`lowerCaseInput` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[11], lowerCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);


const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
