---
id: 68e4276483b0c71658aab1d0
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

異なる大文字・小文字の文字列を新しく作成する必要がある場合があります。

すべての文字を大文字に変換した新しい文字列を作成するには、`toUpperCase()` メソッドを使うことができます。

例を示します。

```js
const sentence = "I love to code!";
sentence.toUpperCase(); // "I LOVE TO CODE!"
```

`upperCaseInput` という変数を作成し、`toUpperCase()` 変数に対して `cleanedInput` メソッドを使った結果を代入してください。

# --hints--

`upperCaseInput` という変数を用意してください。

```js
assert.exists(upperCaseInput);
```

`upperCaseInput` 変数は文字列である必要があります。

```js
assert.isString(upperCaseInput);
```

`cleanedInput.toUpperCase()` を `upperCaseInput` 変数に代入してください。

```js
assert.equal(upperCaseInput, cleanedInput.toUpperCase());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

--fcc-editable-region--

--fcc-editable-region--
```
