---
id: 68e4276483b0c71658aab1d3
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

これまでのレッスンやワークショップで学んだように、`slice()` メソッドを使って文字列の一部を取り出すことができます。

ここで復習しましょう：

```js
const str = "abcd";
str.slice(0, 3); // "abc"
```

空の文字列を `camelCasedVersion` に置き換えて、既存の `lowercaseWord.slice(0, 5)` への代入を変更してください。

そうすれば、コンソールに `"camel"` という単語が表示されるはずです。

# --hints--

`camelCasedVersion` 変数は空の文字列であってはいけません。

```js
assert.isNotEmpty(camelCasedVersion);
```

`lowercaseWord.slice(0, 5)` 変数に `camelCasedVersion` を代入してください。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5));
```

`"camel"` に文字列リテラルの `camelCasedVersion` を代入してはいけません。正しい方法については指示をもう一度読み返してください。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camel(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = ""
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
