---
id: 68e4276483b0c71658aab1cd
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

結果を確認する時間です。

まず、文字列 `console.log()` を使って `"Result of trimming whitespace from both ends:"` を追加してください。

次に、その `console.log()` の下に、`console.log()` 変数を使った別の `cleanedInput` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"Result of trimming whitespace from both ends:"` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[2], "Result of trimming whitespace from both ends:");
```

`cleanedInput` 変数をコンソールにログ出力してください。

```js
assert.equal(getLogs()[3], cleanedInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
--fcc-editable-region--

--fcc-editable-region--
```
