---
id: 68e4276483b0c71658aab1d2
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

新しい文字列を作成するために複数のメソッドを使う必要がある場合があります。次のいくつかのステップで、これまでに学んだメソッドを使ってキャメルケースの命名規則を用いた新しい文字列を作成します。

まず、`lowercaseWord` という変数を作成し、文字列 `"camelcase"` を代入してください。

その `lowercaseWord` 変数の下に、`camelCasedVersion` という別の変数を作成し、今は空の文字列を代入してください。

その `camelCasedVersion` 変数の下に、文字列 `console.log()` を持つ `"Camel cased version:"` を追加し、続けて `console.log()` 変数のための別の `camelCasedVersion` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`lowercaseWord` という変数があるはずです。

```js
assert.exists(lowercaseWord);
```

`lowercaseWord` は文字列であるべきです。

```js
assert.isString(lowercaseWord);
```

`"camelcase"` 変数に文字列 `lowercaseWord` を代入してください。

```js
assert.equal(lowercaseWord, "camelcase");
```

`camelCasedVersion` という変数があるはずです。

```js
assert.exists(camelCasedVersion);
```

`camelCasedVersion` 変数は文字列であるべきです。

```js
assert.isString(camelCasedVersion);
```

`camelCasedVersion` 変数に空の文字列を代入してください。

```js
assert.isEmpty(camelCasedVersion);
```

文字列 `"Camel cased version:"` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[12], "Camel cased version:");
```

`camelCasedVersion` 変数をコンソールにログ出力してください。

```js
assert.equal(getLogs()[13], camelCasedVersion);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

--fcc-editable-region--

--fcc-editable-region--
```
