---
id: 68ed6079fe0e3e98b94a1b75
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

次に、文字列 `console.log()` を持つ `"Original input:"` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"Original input:"` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[0], "Original input:");
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
--fcc-editable-region--

--fcc-editable-region--
```
