---
id: 68f1179eaea3bb59ba434386
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

文字列の末尾の空白を取り除くには、`trimEnd()` メソッドを使うことができます。

例を示します：

```js
const sentence = "  I love learning.  ";
sentence.trimEnd(); // "  I love learning."
```

`trimmedEnd` という変数を作成し、`trimEnd()` 変数に対して `userInput` メソッドを使った結果を代入してください。

# --hints--

`trimmedEnd` という変数があるはずです。

```js
assert.exists(trimmedEnd);
```

`trimmedEnd` 変数は文字列であるべきです。

```js
assert.isString(trimmedEnd);
```

`userInput.trimEnd()` を `trimmedEnd` 変数に代入してください。

```js
assert.equal(trimmedEnd, userInput.trimEnd());
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

--fcc-editable-region--

--fcc-editable-region--
```
