---
id: 68f3eeaa65614107ecd5f1f2
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

最初の `console.log()` の下に、もう一つ `console.log()` を追加して、`userInput` 変数をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

`userInput` 変数をコンソールに出力する必要があります。

```js
assert.equal(getLogs()[1], userInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
--fcc-editable-region--

--fcc-editable-region--
```
