---
id: 68f3f890010c870e17bbe754
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

結果を確認する時間です。

`upperCaseInput` 変数の下に、文字列 `console.log()` を持つ `"Result of using the toUpperCase() method:"` を追加してください。

その `console.log()` の下に、`console.log()` 変数を使った別の `upperCaseInput` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"Result of using the toUpperCase() method:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[8], "Result of using the toUpperCase() method:");
```

`upperCaseInput` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[9], upperCaseInput);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
--fcc-editable-region--

--fcc-editable-region--
```
