---
id: 68f3f3d65d51d80ac0b412fc
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

文字列の先頭の空白をトリムした結果を確認する時間です。

`trimmedStart` 変数の下に、文字列 `console.log()` を持つ `"After using the trimStart() method, leading spaces removed:"` を追加してください。

その下の `console.log()` の下に、`console.log()` 変数を持つ別の `trimmedStart` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"After using the trimStart() method, leading spaces removed:"` をコンソールにログ出力してください。

```js
assert.equal(getLogs()[4], "After using the trimStart() method, leading spaces removed:");
```

`trimmedStart` 変数をコンソールにログ出力してください。

```js
assert.equal(getLogs()[5], trimmedStart);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
--fcc-editable-region--

--fcc-editable-region--
```
