---
id: 68f3f5ca1d2a490c97f14f03
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

文字列の末尾の空白を取り除いた結果を確認する時間です。

`trimmedEnd` 変数の下に、文字列 `console.log()` を持つ `"After using the trimEnd() method, trailing spaces removed:"` を追加してください。

その下の `console.log()` の下に、`console.log()` 変数を持つ別の `trimmedEnd` を追加してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

文字列 `"After using the trimEnd() method, trailing spaces removed:"` をコンソールに出力してください。

```js
assert.equal(getLogs()[6], "After using the trimEnd() method, trailing spaces removed:");
```

`trimmedEnd` 変数をコンソールに出力してください。

```js
assert.equal(getLogs()[7], trimmedEnd);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
--fcc-editable-region--

--fcc-editable-region--
```
