---
id: 68f405c8c9df44132c8fe1bb
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

これまでのレッスンで学んだように、キャメルケースの命名慣例は最初の単語を小文字で始め、続く各単語は先頭が大文字でその後は小文字で続きます。

いくつかの例を思い出してください：

```md
freeCodeCamp
camelCase
learningIsFun
```

`lowercaseWord` 変数の2番目の単語は `"case"` です。その単語の `c` にアクセスするには、`lowercaseWord[5]` を使うことができます。

`camelCasedVersion` の代入を更新して、`+` 演算子を使い、`lowercaseWord.slice(0, 5)` と `lowercaseWord[5]` に対して文字列を大文字に変換する正しいメソッドの結果を連結してください。

これでコンソールに `camelC` が表示されるはずです。

# --hints--

`lowercaseWord.slice(0, 5)` と `lowercaseWord[5].toUpperCase()` を連結してください。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase());
```

リテラル文字列 `"camelC"` を `camelCasedVersion` に代入してはいけません。正しい方法については指示をもう一度読み返してください。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelC(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5)
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```
