---
id: 68f40925d810ae14b2d5a169
title: ステップ 16
challengeType: 1
dashedName: step-16
---

# --description--

現在、`camelCasedVersion` 変数には `camelC` の値が入っています。しかし、望ましい結果は `camelCase` です。

残りの `ase` 文字を取得するには、負のインデックス値を使って `slice()` メソッドを利用できます。負のインデックス値を使うと、`slice()` メソッドが末尾から抽出するように指示されることを思い出してください。

こちらは復習用の例です：

```js
const str = "abcd";
str.slice(-2); // cd
```

`camelCasedVersion` の代入を更新して、`lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` と `slice()` に `lowercaseWord` メソッドを使った結果を連結してください。`-3` メソッドには数値の `slice()` を渡します。

この最後の変更で、ワークショップは完了です！

# --hints--

`lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()` と `lowercaseWord.slice(-3)` を連結してください。

```js
assert.equal(camelCasedVersion, lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase() + lowercaseWord.slice(-3));
```

`"camelCase"` に文字列リテラル `camelCasedVersion` を代入してはいけません。正しい方法については指示をもう一度読んでください。

```js
assert.notMatch(code, /camelCasedVersion\s*=(["'])camelCase(["'])/);
```

# --seed--

## --seed-contents--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";
--fcc-editable-region--
const camelCasedVersion = lowercaseWord.slice(0, 5) + lowercaseWord[5].toUpperCase()
--fcc-editable-region--
console.log("Camel cased version:");
console.log(camelCasedVersion);
```

# --solutions--

```js
const userInput = "   Hello World!   ";
console.log("Original input:");
console.log(userInput);

const cleanedInput = userInput.trim();
console.log("Result of trimming whitespace from both ends:");
console.log(cleanedInput);

const trimmedStart = userInput.trimStart();
console.log("After using the trimStart() method, leading spaces removed:");
console.log(trimmedStart);

const trimmedEnd = userInput.trimEnd();
console.log("After using the trimEnd() method, trailing spaces removed:");
console.log(trimmedEnd);

const upperCaseInput = cleanedInput.toUpperCase();
console.log("Result of using the toUpperCase() method:");
console.log(upperCaseInput);

const lowerCaseInput = cleanedInput.toLowerCase();
console.log("Result of using the toLowerCase() method:");
console.log(lowerCaseInput);

const lowercaseWord = "camelcase";

const camelCasedVersion =
  lowercaseWord.slice(0, 5) +
  lowercaseWord[5].toUpperCase() +
  lowercaseWord.slice(-3);

console.log("Camel cased version:");
console.log(camelCasedVersion);
```
