---
id: 6852f11d412c29665985a915
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Тепер, коли користувачі можуть надсилати е-пошту, їм потрібен спосіб зберігати отримані листи. Кожен користувач має мати вхідну скриньку для збору своїх листів.

Додайте атрибут `inbox` до методу `User` класу `__init__` і ініціалізуйте його як порожній список `[]`. Він зберігатиме всі листи, які отримує користувач.

# --hints--

Ви маєте додати `self.inbox = []` до методу `User` класу `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
