---
id: 6852f06fdde1b4634a263f03
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Тепер давайте перевіримо клас `Email`, створивши об’єкт е-пошти та перевіривши, що він правильно зберігає інформацію. Ви виведете кілька атрибутів, щоб переконатися, що все працює.

Створіть об’єкт е-пошти з іменем `email_obj` з `alice@example.com` як відправником, `bob@example.com` як отримувачем, `Hello` як темою та `Hi Bob!` як тілом листа. Потім виведіть атрибути `sender` і `subject` окремими інструкціями виведення, щоб перевірити, що вони збережені правильно.

# --hints--

Ви повинні створити об’єкт е-пошти з іменем `email_obj` з вказаними параметрами.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

Ви повинні вивести `email_obj.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

Ви повинні вивести `email_obj.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
