---
id: 6852f0c3fcedb964ab959921
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Тепер давайте перевіримо, що атрибут `read` було додано правильно до вашого класу `Email`. Оскільки у вас уже є об’єкт е-пошти з попередніх кроків, виведіть атрибут `read`, щоб переконатися, що він тепер `False` за замовчуванням.

# --hints--

Ви маєте вивести `email_obj.read`, щоб перевірити атрибут для читання.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
