---
id: 6852ea0c3cadb54e897acdac
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите *симулятор е-пошти*, який імітує надсилання, отримання та керування е-поштою між різними користувачами. Ви дізнаєтеся про класи, об’єкти та як організувати код у об’єктно-орієнтований спосіб.

Почніть зі створення класу з назвою `Email` за допомогою ключового слова `class`.

# --hints--

Ви маєте створити клас з назвою `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
