---
id: 6852ef95ccc27660275692cd
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Ваш клас `Email` має зберігати інформацію про кожну е-пошту під час її створення.

Всередині вашого класу `Email` видаліть існуюче ключове слово `pass` і створіть метод `__init__`, який прийматиме параметри `self`, `sender`, `receiver`, `subject` і `body`.

# --hints--

Ви повинні визначити метод `__init__` всередині класу `Email`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

Ваш метод `__init__` має мати параметри: `self`, `sender`, `receiver`, `subject` і `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
