---
id: 6852f01116912861c7d1b0e6
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

У методі `__init__` потрібно призначити значення параметрів атрибутам об’єкта, щоб кожен email міг зберігати свою інформацію.

У вашому методі `__init__` видаліть ключове слово `pass` і призначте параметри `self.sender`, `self.receiver`, `self.subject` та `self.body`.

# --hints--

Ви повинні призначити параметр `sender` атрибуту `self.sender`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

Ви повинні призначити параметр `receiver` атрибуту `self.receiver`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

Ви повинні призначити параметр `subject` атрибуту `self.subject`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

Ви повинні призначити параметр `body` атрибуту `self.body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
