---
id: 6853f4584f508e95cc7a9f4d
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Тепер ви створите метод для виведення списку всіх е-поштових повідомлень у вхідних. Цей метод має обробляти випадок, коли вхідні порожні, і відображати нумерований список е-поштових повідомлень, якщо вони є.

Додайте метод з назвою `list_emails` до вашого класу `Inbox`, який прийматиме `self` як параметр. Спочатку створіть інструкцію `if` для перевірки, чи порожні вхідні, тестуючи список `emails`. Якщо він порожній, виведіть повідомлення `Your inbox is empty.\n`. Додайте інструкцію `return` для виходу з методу.

# --hints--

Ви маєте визначити метод з назвою `list_emails` у класі `Inbox`, який прийматиме `self` як параметр.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").has_args("self")`))
    }
})
```

Ви маєте перевірити, чи `self.emails` порожній, використовуючи `if not self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})
```

Ви маєте вивести повідомлення `Your inbox is empty.\n`, якщо немає е-поштових повідомлень.

```js
({
    test: () => 
    { 
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[0].is_equivalent("print('Your inbox is empty.\\\\n')")`))
    }
})
```

Ви маєте мати інструкцію `return` в кінці інструкції `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_ifs()[0].find_body()[1].is_equivalent("return")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    
--fcc-editable-region--

```
