---
id: 6853f4ded9f84896ff182dd8
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Після перевірки порожньої папки вхідних виведіть повідомлення `\nYour Emails:`. Потім ітеруйте над усіма листами у папці вхідних за допомогою циклу `for` з нумерацією. Ось синтаксис для нумерації з початковим числом: `enumerate(iterable, start=start_number)`.

Використовуйте нумерацію, щоб пронумерувати листи, починаючи з `1`. Використовуйте `i` (індекс) і `email` (об’єкт листа) як змінні ітерації.

Усередині циклу виведіть f-рядок з індексом ітерації, за яким іде `.`, потім пробіл і рядкове представлення об’єкта листа у такому форматі:

```md
i. string_representation
```

Де `i` — це індекс, а `string_representation` — представлення об’єкта листа.

# --hints--

Ви повинні вивести заголовок `\nYour Emails:` поза інструкцією `if`.

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_body()[1].is_equivalent("print('\\\\nYour Emails:')")`))
    }
})
```

У вас має бути цикл `for`, який ітерує над `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        runPython(`
_iter = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_iter()
assert _iter.is_equivalent("enumerate(self.emails, start=1)") or _iter.is_equivalent("enumerate(self.emails, 1)")
`)
    }
})
```

Ваш цикл `for` має використовувати `i` і `email` для ітерації над `enumerate(self.emails, start=1)`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_for_vars().is_equivalent("i, email")`))
    }
})
```

Усередині циклу ви маєте виводити кожне резюме листа з пронумерованим префіксом, за яким іде форматований рядок із методу `__str__`. Не забудьте додати `.` після числа.

```js
({
    test: () => 
    {
        assert(runPython(`
        cond = _Node(_code).find_class("Inbox").find_function("list_emails").find_for_loops()[0].find_body()
        cond.is_equivalent("print(f'{i}. {str(email)}')") or cond.is_equivalent("print(f'{i}. {email}')")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        
--fcc-editable-region--
```
