---
id: 6859400f8f7f872e50e1238c
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Клас `User` потребує методу для надсилання е-пошти іншим користувачам. Під час надсилання е-пошти ви створите новий об’єкт `Email` і додасте його до вхідних повідомлень отримувача.

Створіть метод з назвою `send_email` у вашому класі `User`. Цей метод має приймати параметри для `receiver`, `subject` та `body`.

Всередині методу створіть новий об’єкт `Email` із такими значеннями параметрів і призначте його змінній з ім’ям `email`:

- `sender`: `self` (поточний користувач)
- `receiver`: `receiver` (користувач, який отримує е-пошту)
- `subject`: `subject` (тема е-пошти)
- `body`: `body` (текст е-пошти)

# --hints--

Ви маєте визначити метод `send_email` у класі `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

Метод `send_email` має мати параметри: `self`, `receiver`, `subject` та `body`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

Метод має створити об’єкт `Email` з ім’ям `email` з правильними параметрами.

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
