---
id: 68594df7779a164e6a61e8ae
title: Крок 32
challengeType: 20
dashedName: step-32
---

# --description--

Потрібно додати метод для вхідних повідомлень, який зчитуватиме конкретний лист. Коли користувач хоче прочитати лист, він вказує номер листа (починаючи з індексу 0), який хоче побачити, і метод виводить повний вміст листа.

Додайте метод з назвою `read_email` до вашого класу `Inbox`, який прийматиме параметр `index`. Спочатку перевірте, чи порожня вхідна скринька, і виведіть повідомлення `Inbox is empty.\n`, якщо це так. Додайте інструкцію `return` після цього, щоб вийти з методу.

# --hints--

Ви повинні визначити метод з назвою `read_email` у класі `Inbox`, який прийматиме параметри `self` і `index`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").has_args("self, index")`));
    }
})
```

Ви повинні перевірити, чи `self.emails` порожня, використовуючи `if not self.emails`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`));
    }
})
```

Ви повинні вивести `Inbox is empty.\n`, якщо листів немає, і повернутися.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[0].is_equivalent("print('Inbox is empty.\\\\n')")`));
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_ifs()[0].find_body()[1].is_equivalent("return")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    
--fcc-editable-region--
```
