---
id: 68593ea22eac7e228ed585ce
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Е-пошта має відстежувати, чи її прочитано. Додайте атрибут `read` до методу `__init__` і призначте йому значення `False` за замовчуванням, оскільки нові листи починаються як непрочитані.

# --hints--

Ви маєте призначити `False` для `self.read` у методі `__init__`.

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
