---
id: 685cdabd4cded3b340193f72
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Тепер додайте рядок теми до відображення е-пошти. У методі `display_full_email` додайте інструкцію виведення, щоб показати тему е-пошти у такому форматі:

- `Subject: subject`, де `subject` замінений на тему е-пошти

# --hints--

Ви повинні використовувати f-рядки, щоб вивести тему, використовуючи `Subject: {self.subject}`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'Subject: {self.subject}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
