---
id: 685c1eba966c6396850736ea
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Хоч користувачі можуть надсилати е-пошту та мати скриньки, нам потрібен спеціальний клас для ефективного керування операціями зі скринькою.

Клас `Inbox` зберігатиме список е-пошти та надаватиме методи для додавання нових листів, переліку всіх листів і їхнього керування.

Створіть новий клас із назвою `Inbox` з методом `__init__`, який ініціалізує порожній список під назвою `emails`.

# --hints--

Ви маєте створити клас `Inbox`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

Клас `Inbox` повинен мати метод `__init__`, який приймає `self` як параметр.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

Ви маєте ініціалізувати `self.emails = []` у методі `Inbox` класу `__init__`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
