---
id: 685c1e32254e5993a5c8823e
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер, коли у вас є ініціальне налаштування класу `Email`, настав час створити користувачів, які можуть надсилати та отримувати е-пошту. Кожен користувач матиме ім’я і зможе виконувати операції з е-поштою.

Створіть клас `User` з методом `__init__`, який приймає `self` та `name` як параметри. Усередині класу призначте параметр `name` до `self.name`.

# --hints--

Ви повинні створити клас `User`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

Клас `User` має мати метод `__init__`, який приймає `self` та `name` як параметри.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

Ви повинні призначити параметр `name` до `self.name`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
