---
id: 685d2fd1f309b9191f257c55
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Перед тим як інтегрувати часові позначки у нашу систему е-пошти, давайте потренуємося працювати з форматуванням datetime. Функція `datetime.datetime.now()` дає нам поточну дату і час, а ми можемо використати метод `strftime()`, щоб форматувати їх по-різному.

Ось як `strftime()` працює з кодами форматування:

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

Коди форматування, як `%Y` (рік), `%m` (місяць), `%d` (день), кажуть `strftime()`, що включати, і ви можете додавати роздільники, як `-`, між ними.

Внизу вашого коду створіть змінну з назвою `current_time` і призначте їй `datetime.datetime.now()`. Потім використайте `strftime()`, щоб вивести час у форматі години:хвилини:секунди, використовуючи `:` як роздільник.

Використайте ці коди форматування: `%H` для годин (24-годинний формат), `%M` для хвилин і `%S` для секунд.

# --hints--

Ви повинні створити змінну з назвою `current_time` і призначити їй `datetime.datetime.now()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

Ви повинні використати `strftime()`, щоб вивести час у форматі `%H:%M:%S`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
