---
id: 68875554c6c32de788afc64e
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Користувачам потрібен спосіб правильно читати свої е-пошти.

Додайте метод `display_full_email` до класу `Email`, який прийматиме лише `self` як параметр. Усередині цього методу викличте метод `mark_as_read()`, щоб позначити е-пошту як прочитану, коли хтось її переглядає.

Цей метод використовуватиметься для відображення повного вмісту е-пошти у гарно відформатованому вигляді.

# --hints--

Ви повинні визначити метод `display_full_email` у класі `Email`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

Метод `display_full_email` має приймати лише `self` як параметр.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

Метод має викликати `self.mark_as_read()`, щоб позначити е-пошту як прочитану.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
