---
id: 688754d6d59578df1313cb3c
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Тепер настав час зв’язати класи User і Inbox, щоб електронні листи могли справді доставлятися!

Користувачі повинні мати відповідні об’єкти Inbox замість простих списків, а метод `send_email` має доставляти листи до поштової скриньки отримувача.

Оновіть клас `User` `__init__` так, щоб кожен користувач отримував справжній об’єкт Inbox замість порожнього списку, створеного раніше.

# --hints--

Вам слід оновити клас User, щоб ініціалізувати `self.inbox = Inbox()`.

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
