---
id: 68b9350183ff05bd106c79e4
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Тепер протестуйте метод `mark_as_read`, який ви створили в класі `Email`. Використайте цей метод на існуючому об’єкті `Email` `email_obj`, щоб змінити статус. Після позначення е-пошти як прочитаної, виведіть атрибут `read` об’єкта `email_obj`, щоб підтвердити зміну.

# --hints--

Ви повинні викликати метод `mark_as_read` на `email_obj`, щоб позначити е-пошту як прочитану.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

Після виклику методу `read` ви повинні вивести атрибут `email_obj` об’єкта `mark_as_read`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
