---
id: 68b9366fae8b13c1278761b1
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Вашій поштовій скриньці потрібен спосіб отримувати нові листи. Коли хтось надсилає лист користувачу, його слід додати до його поштової скриньки.

Додайте метод із назвою `receive_email` до вашого класу `Inbox`, який прийматиме `self` та об’єкт листа `email` як параметри. Усередині тіла методу додайте `email` до списку `emails` за допомогою методу `append`.

# --hints--

Ви маєте визначити метод із назвою `receive_email` у класі `Inbox`, який прийматиме `self` та `email` як параметри.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

Ви маєте додати `email` до `self.emails`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
