---
id: 68b930ee0e7b8fb81895e59a
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Тепер ви додасте метод для позначення е-пошти як прочитаної. Коли хтось відкриває е-пошту, ви хочете змінити її статус з непрочитаної на прочитану. Цей метод буде простим — він лише має встановити атрибут `read` у значення `True`.

Додайте метод з назвою `mark_as_read` до вашого класу `Email`. Цей метод повинен приймати лише параметр `self` і встановлювати атрибут `read` у значення `True`.

# --hints--

Ви маєте визначити метод з назвою `mark_as_read` у класі `Email`, який прийматиме лише параметр `self`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

Метод має встановити `self.read` у значення `True`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
