---
id: 68be6a4447e61c4395ef2359
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

Вашій поштовій скриньці також потрібен спосіб видаляти листи.

Ключове слово `del` у Python можна використати, щоб видалити елементи зі списку за їхнім індексом.

Додайте метод із назвою `delete_email` до свого класу `Inbox`. Як і `read_email`, він має:

- Приймати параметр `index`, перевіряти, чи порожня пошта, і виводити повідомлення `Inbox is empty.\n`, якщо це так. Використайте `return`, щоб вийти з методу.

# --hints--

Ви маєте визначити метод із назвою `delete_email` у класі `Inbox`, який приймає `self` та `index` як параметри.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").has_args("self,index")`))
    }
})

```

Спочатку слід перевірити, чи порожня пошта, використовуючи `if not self.emails:`.

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_conditions()[0].is_equivalent("not self.emails")`))
    }
})

```

Якщо пошта порожня, слід вивести повідомлення `Inbox is empty.\n` і використати `return`, щоб вийти з методу.

```js
({
    test: () => 
    {
    assert(runPython(`_Node(_code).find_class("Inbox").find_function("delete_email").find_ifs()[0].find_body().is_equivalent("print('Inbox is empty.\\\\n')\\nreturn")`))
}
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()
--fcc-editable-region--

--fcc-editable-region--
```
