---
id: 6823ac607bfdbc46331b2559
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

HTML تعني لغة الترميز النصي الفائق. هو الكود الذي يحدد هيكل ومحتوى صفحة الويب. هذا هو محرر الكود الخاص بك، حيث ستكتب HTML.

ابحث عن السطر 1 في المحرر واكتب هذا النص:

```md
Welcome to freeCodeCamp
```

عندما تنتهي، انقر على زر "التحقق من الكود" لترى ما إذا كان صحيحًا.

# --hints--

يجب أن يحتوي محررك على النص `Welcome to freeCodeCamp`. تحقق من الإملاء جيدًا.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /welcome\s*to\s*freecodecamp/i)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
