---
id: 6823d9ac8bdc3853df65a1ff
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

هناك ستة عناصر عنوان في HTML: من `h1` إلى `h6`. تُستخدم لإظهار أهمية الأقسام في صفحة الويب الخاصة بك، حيث يكون `h1` هو الأهم و`h6` الأقل أهمية.

أسفل عنصر `p` الخاص بك، أضف عنوان `h3` بالنص التالي:

```md
Introduction to HTML
```

# --hints--

يجب أن يحتوي عنصر `h3` الخاص بك على وسم افتتاحي `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

يجب أن يحتوي عنصر `h3` الخاص بك على وسم إغلاق `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

يجب أن يكون نص عنصر `h3` الخاص بك هو `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
