---
id: 6823d6244511f252c8300eed
title: الخطوة 6
challengeType: 0
dashedName: step-6
---

# --description--

عندما تحتاج إلى إضافة فقرة إلى صفحة ويب، يمكنك استخدام العنصر `p` بهذه الطريقة:

```html
<p>This is a paragraph element.</p>
```

حوّل النص `Learn the skills to become a full-stack developer` إلى عنصر فقرة.

# --hints--

يجب أن يحتوي عنصر `p` على وسم بداية. صيغة وسم البداية هي: `<elementName>`.

```js
assert.exists(document.querySelector("p"));
```

يجب أن يحتوي عنصر `p` على وسم إغلاق. صيغة وسم الإغلاق هي: `</elementName>`.

```js
assert.match(code, /<\/p\s*\>/);
```

يجب أن يكون نص عنصر `p` هو `Learn the skills to become a full-stack developer`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*Learn\s*the\s*skills\s*to\s*become\s*a\s*full-stack\s*developer\s*\<\/p\s*\>/i);
```

يجب أن يكون عنصر `p` أسفل عنصر `h2`.

```js
assert.exists(document.querySelector("h2 + p"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
--fcc-editable-region--
Learn the skills to become a full-stack developer
--fcc-editable-region--
```
