---
id: 6823e036ea4b71553558c01b
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

لاحظ أن كل عنوان يبدو مختلفًا قليلاً في المعاينة. نفس الشيء مع الفقرة.

بعد ذلك، أضف عنصر `p` آخر يعرض النص التالي:

```md
HTML represents the content and structure of a webpage
```

# --hints--

يجب أن يكون لديك عنصر `p` ثاني.

```js
assert.lengthOf(document.querySelectorAll("p"), 2);
```

يجب أن يحتوي عنصر `p` الخاص بك على النص `HTML represents the content and structure of a webpage`. يمكنك نسخ النص من التعليمات.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*HTML\s*represents\s*the\s*content\s*and\s*structure\s*of\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
--fcc-editable-region--

--fcc-editable-region--
```
