---
id: 6823e169fda14755fbf00445
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

أنت تتقن الأمر. بعد ذلك، أضف عنصر `h3` آخر في أسفل المحرر مع النص:

```md
Introduction to CSS
```

# --hints--

يجب أن تضيف عنصر `h3` ثاني إلى الصفحة.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

يجب أن يكون نص عنصر `h3` هو `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```
