---
id: 6823e637c1c0ed56f781b4fc
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

أخيرًا، تجعل JavaScript صفحة الويب الخاصة بك تفاعلية — تتيح لك إخبار الصفحة بما يجب فعله عندما ينقر شخص ما على زر، أو يرسل نموذجًا، أو أشياء كثيرة أخرى.

لآخر خطوة في ورشة العمل، أضف عنصر `h3` وعنصر `p` آخرين إلى الصفحة يصفان JavaScript.

أولًا، أضف عنصر `h3` بالنص التالي:

```md
Introduction to JavaScript
```

ثم، أسفل ذلك العنصر `h3`، أضف عنصر `p` بالنص التالي:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

يجب أن تضيف عنصر `h3` ثالث إلى الصفحة.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

يجب أن يكون نص عنصر `h3` الخاص بك هو `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

يجب أن تضيف عنصر `p` رابع إلى الصفحة.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

يجب أن يحتوي عنصر `p` الخاص بك على النص `JavaScript adds interactivity to a webpage`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
