---
id: 682ba2318000b62f179bdf04
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

تتكون HTML من عناصر. أول عنصر ستستخدمه هو العنصر `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

يبدأ بوسم البداية (`<h1>`)، وينتهي بوسم الإغلاق (`</h1>`)، ويحتوي على النص الذي سيعرضه بين الوسمين.

حوّل نص `Welcome to freeCodeCamp` إلى عنصر `h1` بإضافة وسم البداية قبله، ووسم الإغلاق بعده.

# --hints--

يجب أن يحتوي عنصر `h1` على وسم بداية `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

يجب أن يحتوي عنصر `h1` على وسم إغلاق `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

يجب أن يبدو عنصر `h1` هكذا: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
