---
id: 682cd206883fc7b25eb539c5
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

لاحظ أن HTML الذي تكتبه في المحرر يظهر في المعاينة. في هذه الورشة، ستكتب HTML لصفحة ويب جزئية للمناهج.

أسفل عنصر `h1` الخاص بك، اكتب التالي في السطر الفارغ:

```md
Full-Stack Curriculum
```

# --hints--

يجب أن يحتوي محررك على النص `Full-Stack Curriculum`. تحقق من صحة التهجئة.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
