---
id: 683921c4769fd23dbadec2fe
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

بينما يحدد HTML هيكل ومحتوى صفحة الويب، يُستخدم CSS لإضافة التنسيق — أشياء مثل الألوان، الخطوط، التباعد، والتخطيط.

أسفل العنصر `h3` الذي أضفته للتو، أضف عنصر فقرة آخر بالنص التالي:

```md
CSS is used to style a webpage
```

# --hints--

يجب أن يكون لديك بالضبط عنصران `h3`، تأكد من عدم إزالة العناصر الموجودة، وعدم إضافة عناصر جديدة.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

لا يجب تغيير نص العنصر الثاني `h3`، يجب أن يظل مكتوبًا `Introduction to CSS`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

يجب أن تضيف عنصر `p` ثالث إلى الصفحة.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

يجب أن يحتوي عنصر `p` الخاص بك على النص `CSS is used to style a webpage`. يمكنك نسخ النص من التعليمات.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
--fcc-editable-region--

--fcc-editable-region--
```
