---
id: 67c562286b29447da020d407
title: Erstellen Sie einen Einmal-Passwort-Generator
challengeType: 25
dashedName: build-a-one-time-password-generator
demoType: onClick
---

# --description--

In diesem Labor erzeugen Sie ein 6-stelliges OTP (One-Time Password) und zeigen es dem Benutzer an. Das OTP läuft nach 5 Sekunden ab, und ein neues OTP wird generiert, wenn der Benutzer auf die Schaltfläche `"Generate OTP"` klickt.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten den `useEffect`-Hook verwenden, um den Countdown-Timer zu verwalten.
2. Ihre `OTPGenerator`-Komponente sollte ein `div`-Element mit dem Klassennamen `container` zurückgeben.
3. Ihr `div` mit der Klasse `container` sollte die folgenden Elemente enthalten:

    - Ein `h1`-Element mit der ID `otp-title` und dem Text `"OTP Generator"`.
    - Ein `h2`-Element mit der ID `otp-display`, das entweder die Nachricht `"Click 'Generate OTP' to get a code"` anzeigt oder das generierte OTP, falls eines vorhanden ist.
    - Ein `p`-Element mit der ID `otp-timer` und einem `aria-live`-Attribut, das auf einen gültigen Wert gesetzt ist und:
        - Anfangs leer ist.
        - Nach dem Klicken auf die Schaltfläche die Meldung `"Expires in: X seconds"` anzeigt, wobei `X` die verbleibende Zeit bis zum Ablauf des OTP darstellt.
        - Die Nachricht `"OTP expired. Click the button to generate a new OTP."` anzeigt, sobald der Countdown `0` erreicht.
    - Ein `button`-Element mit der ID `generate-otp-button` und der Bezeichnung `"Generate OTP"`. Beim Klicken soll ein neues OTP generiert und ein 5-Sekunden-Countdown gestartet werden.
    - Die Schaltfläche `"Generate OTP"` muss deaktiviert sein, solange der Countdown aktiv ist.

4. Sie sollten sicherstellen, dass der Countdown-Timer automatisch stoppt, sobald er `0` Sekunden erreicht, um unnötige Aktualisierungen zu vermeiden.
5. Ihr generiertes OTP sollte 6 Stellen lang sein.

# --before-all--

```js
let clock = __FakeTimers.install();
```

# --after-all--

```js
clock.uninstall();
```

# --hints--

Sie sollten die bestehende `const { useState, useEffect, useRef } = React;` Zuweisung im Code nicht entfernen.

```js
assert.match(code, /{\s*(useState\s*,\s*useEffect\s*,\s*useRef|useState\s*,\s*useRef\s*,\s*useEffect|useEffect\s*,\s*useState\s*,\s*useRef|useEffect\s*,\s*useRef\s*,\s*useState|useRef\s*,\s*useState\s*,\s*useEffect|useRef\s*,\s*useEffect\s*,\s*useState|useState|useEffect|useRef)\s*}\s*=\s*React;/);
```

Sie sollten ein `div` mit der Klasse `.container` rendern.

```js
const containerDiv = document.querySelector('.container');
assert.exists(containerDiv);
```

Ihr `.container` sollte ein `h1`-Element mit der ID `otp-title` und dem Text `"OTP Generator"` enthalten.

```js
const h1 = document.querySelector(".container h1#otp-title");
assert.exists(h1);
assert.equal(h1.textContent, "OTP Generator");
```

Ihr `.container` sollte ein `h2`-Element mit der ID `otp-display` enthalten, das das OTP anzeigt, wenn es generiert wurde.

```js
const h2 = document.querySelector(".container h2#otp-display");
assert.exists(h2);
```

Ihr `h2` innerhalb von `.container` sollte anfangs die Nachricht `"Click 'Generate OTP' to get a code"` anzeigen.

```js
assert.strictEqual(document.querySelector('.container h2#otp-display').textContent.trim(), "Click 'Generate OTP' to get a code");
```

Ihr `.container`-Element sollte ein `p`-Element mit der ID `otp-timer` enthalten.

```js
const container = document.querySelector('.container');
assert.exists(container, "The container element should exist");
const pElements = container.querySelectorAll('p');
assert.strictEqual(pElements.length, 1);
assert.strictEqual(pElements[0].id, "otp-timer", "The p element should have the id otp-timer");
```

Ihr `p`-Element sollte ein `aria-live`-Attribut mit dem Wert `assertive` oder `polite` haben.

```js
const pElement = document.querySelector('.container p');
assert.exists(pElement?.ariaLive);
assert.oneOf(pElement?.ariaLive, ["assertive", "polite"]);
```

Ihr `p`-Element sollte anfangs leer sein.

```js
const pElement = document.querySelector('.container p');
assert.strictEqual(pElement?.innerText, "", "The `p` element should be empty");
```

Ihr `.container` sollte ein `button`-Element mit der ID `generate-otp-button` und dem Text `"Generate OTP"` enthalten.

```js
const button = document.querySelector('.container button#generate-otp-button').textContent;
assert.exists(button);
assert.strictEqual(button, 'Generate OTP');
```

Ihr `h2`-Element mit der ID `otp-display` sollte ein neues OTP anzeigen, wenn die Schaltfläche geklickt wird.

```js
try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)

assert.match(h2.textContent?.trim(), /[0-9]/, "OTP should be a number");
} finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Ihr generiertes OTP sollte 6 Stellen lang sein.

```js
  try {
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1);
const h2 = document.querySelector(".container h2#otp-display");

assert.match(h2.textContent.trim(), /^[0-9]{6}$/, "OTP should be a 6-digit number");

 } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```


Ihr `p`-Element mit der ID `otp-timer` sollte beim Klicken auf die Schaltfläche einen 5-Sekunden-Countdown anzeigen.

```js
  try {
// allow the page to render
clock.tick(1)
const button = document.querySelector(".container button#generate-otp-button");
const h2 = document.querySelector(".container h2#otp-display");

button.click();

clock.tick(300)
const p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/, "Countdown should start at 5 seconds");

  } finally {
// we need to wait for the countdown to finish for the next test to work otherwise the button will be disabled
await clock.runAllAsync();
}
```

Ihr `p`-Element mit der ID `otp-timer` sollte seine Meldung jede Sekunde aktualisieren, um die verbleibende Zeit anzuzeigen.

```js
const button = document.querySelector(".container button#generate-otp-button");
button.click();
clock.tick(1)
let p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 5 seconds/);
await clock.tickAsync(1000)
p = document.querySelector(".container p#otp-timer");
assert.match(p.textContent, /Expires in: 4 seconds/, "Countdown should update to 4 seconds after 1 second");
```

Ihre Schaltfläche `"Generate OTP"` sollte deaktiviert sein, solange der Countdown-Timer läuft.

```js
try {
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();
  clock.tick(1);
  
  assert.isTrue(button.disabled);
} finally {
  await clock.runAllAsync();
}
```

Ihre Schaltfläche `"Generate OTP"` sollte aktiviert werden, sobald der Countdown-Timer 0 erreicht.

```js

const button = document.querySelector(".container button#generate-otp-button");
button.click();
await clock.tickAsync(6000);
  
assert.isFalse(button.disabled);
```

Sie sollten die Nachricht `OTP expired. Click the button to generate a new OTP.` anzeigen, wenn der Countdown-Timer `0` erreicht.

```js
  const button = document.querySelector(".container button#generate-otp-button");
  button.click();

  await clock.tickAsync(6000);

  const p = document.querySelector(".container p#otp-timer");
  assert.equal(p.textContent.trim(), "OTP expired. Click the button to generate a new OTP.");
```

Sie sollten die `OTPGenerator` Komponente exportieren.

```js
assert.isFunction(window.index.OTPGenerator);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <link rel="stylesheet" href="styles.css" />
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css

```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {};
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>OTP Generator</title>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
    <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      src="index.jsx"
    ></script>
    <link rel="stylesheet" href="styles.css" />
</head>

<body>
    <div id="root"></div>
    <script
      data-plugins="transform-modules-umd"
      type="text/babel"
      data-presets="react"
      data-type="module"
    >
      import { OTPGenerator } from './index.jsx';
      ReactDOM.createRoot(document.getElementById('root')).render(<OTPGenerator />);
    </script>
</body>

</html>
```

```css
body {
    font-family: 'Castoro Titling', cursive;
}

.container {
    text-align: center;
    margin-top: 50px;
}

h1 {
    font-size: 30px;
    color: #333;
}

h2 {
    font-size: 25px;
    color: #555;
}

p {
    font-size: 16px;
    color: #888;
}

button {
    padding: 10px 20px;
    margin-top: 20px;
    font-size: 16px;
    border: none;
    background-color: #383f59;
    color: white;
    border-radius: 5px;
}

button:disabled {
  background-color:#57585c;
  cursor: not-allowed;
}

button:hover {
    background-color: #444;
}
```

```jsx
const { useState, useEffect, useRef } = React;

export const OTPGenerator = () => {
  const intervalId = useRef(null);
  const [otp, setOtp] = useState(null);
  const [secondsLeft, setSecondsLeft] = useState(5);
  const [isCounting, setIsCounting] = useState(false);
  const [hasGeneratedOtp, setHasGeneratedOtp] = useState(false);

  const generateOTP = () => Math.floor(100000 + Math.random() * 900000);

  const handleGenerateOtp = () => {
    setOtp(generateOTP());
    setSecondsLeft(5);
    setIsCounting(true);
    setHasGeneratedOtp(true);
    intervalId.current = setInterval(() => {
      setSecondsLeft(prevSeconds => prevSeconds - 1);
    }, 1000);
  };

  useEffect(() => {
    if (secondsLeft === 0) {
      setIsCounting(false);
      clearInterval(intervalId.current);
    }
  }, [secondsLeft]);

  useEffect(
    () => () => {
      clearInterval(intervalId.current);
    },
    []
  );

  return (
    <div className='container'>
      <h1 id='otp-title'>OTP Generator</h1>
      <h2 id='otp-display'>
        {otp ? otp : "Click 'Generate OTP' to get a code"}
      </h2>
      <p id='otp-timer' aria-live='assertive'>
        {isCounting
          ? `Expires in: ${secondsLeft} seconds`
          : hasGeneratedOtp && 
            "OTP expired. Click the button to generate a new OTP."}
      </p>
      <button
        id='generate-otp-button'
        onClick={handleGenerateOtp}
        disabled={isCounting}
      >
        Generate OTP
      </button>
    </div>
  );
};
```
