---
id: 68e6698b29e562d21b1d260c
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在這個工作坊中，你將建置一個主要網頁瀏覽器的列表。已經為你提供了 HTML 樣板。

首先在你的頁面中，使用 `body` 元素內的 `h1` 元素添加一個標題，內容為 `List of Major Web Browsers`。

# --hints--

你應該在 `body` 元素內有一個 `h1` 元素。

```js
const h1 = document.querySelector('body h1');
assert.exists(h1);
```

你的 `h1` 元素應該包含文字 `List of Major Web Browsers`。

```js
const h1 = document.querySelector('body h1');
assert.strictEqual(h1?.textContent.trim(), 'List of Major Web Browsers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="UTF-8">
        <meta name="viewport" content="width=device-width, initial-scale=1.0">
        <title>List of Browsers and Descriptions</title>
    </head>
    <body>
    --fcc-editable-region--
        
    --fcc-editable-region--
    </body>
</html>
```
