---
id: 68ecfa1003d88990bb96a33d
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

描述列表由一個術語和詳細資訊的列表組成。`dt` 元素表述描述術語，而 `dd` 元素表述描述詳細資訊。 

你將為你的描述列表添加的第一個瀏覽器將是 Google Chrome。

在你的 `dl` 元素內，為你的 `dt` 元素添加包含文字 `Google Chrome`。

在你的 `dt` 元素下方，添加一個帶有以下文字的 `dd` 元素：

```md
This is a free web browser developed by Google and first released in 2008.
```

# --hints--

你應該在你的 `dl` 元素內有一個 `dt` 元素。

```js
const dlElement = document.querySelector('dl dt');
assert.exists(dlElement);
```

你的 `dt` 元素應該有文字 `Google Chrome`。

```js
const dt = document.querySelector('dt');
assert.strictEqual(dt?.textContent.trim(), "Google Chrome");
```

你應該在你的 `dt` 元素下方有一個 `dd` 元素。

```js
const ddElement = document.querySelector('dl dd');
assert.exists(ddElement);
```

你的 `dd` 元素應包含文字 `This is a free web browser developed by Google and first released in 2008.`

```js
const ddElement = document.querySelector('dd');
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by Google and first released in 2008.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
        --fcc-editable-region--
            
        --fcc-editable-region--
        </dl>
    </body> 
</html>
```

