---
id: 68ecfa0f03d88990bb96a33c
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

如你從先前的課程中回想，描述列表用於以有組織且易於閱讀的格式呈現術語和定義。 

這裡有一個範例：

```html
<dl>
  <dt>HTML</dt>
  <dd>HyperText Markup Language</dd>
  <dt>CSS</dt>
  <dd>Cascading Style Sheets</dd>
</dl>
```

在 `h1` 元素下方，創建一個 `dl` 元素。這將用來存放瀏覽器的列表。

# --hints--

你應該有一個 `dl` 元素。

```js
const dl = document.querySelector('dl');
assert.exists(dl);
```

你的 `dl` 元素應該位於 `h1` 元素下方。

```js
assert.exists(document.querySelector('h1 + dl'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        --fcc-editable-region--
        
        --fcc-editable-region--
    </body> 
</html>
```

