---
id: 68ecfa1003d88990bb96a33e
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

現在是時候將你的第二個瀏覽器添加到列表中。

在你的 `Google Chrome` 項目下方，添加另一個包含文字 `Firefox` 的 `dt` 元素。

在你的 `dt` 元素下方，添加一個帶有以下文字的 `dd` 元素：

```md
This is a free web browser developed by the Mozilla Corporation and first created in 2004.
```

# --hints--

你的程式碼中應該有兩個 `dt` 元素。

```js
const dtElements = document.querySelectorAll('dt');
assert.lengthOf(dtElements, 2);
```

你的第二個 `dt` 元素應該有文字 `Firefox`。

```js
const dt = document.querySelectorAll('dt')[1];
assert.strictEqual(dt?.textContent.trim(), "Firefox");
```

你的程式碼中應該有兩個 `dd` 元素。

```js
const ddElements = document.querySelectorAll('dd');
assert.lengthOf(ddElements, 2);
```

你的第二個 `dd` 元素應該包含文字 `This is a free web browser developed by the Mozilla Corporation and first created in 2004.`

```js
const ddElement = document.querySelectorAll('dd')[1];
assert.strictEqual(ddElement?.textContent.trim(), "This is a free web browser developed by the Mozilla Corporation and first created in 2004.");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html> 
<html lang="en"> 
    <head> 
        <meta charset="UTF-8"> 
        <meta name="viewport" content="width=device-width, initial-scale=1.0"> 
        <title>List of Browsers and Descriptions</title> 
    </head> 
    <body> 
        <h1>List of Major Web Browsers</h1> 
        <dl>
            <dt>Google Chrome</dt>
            <dd>This is a free web browser developed by Google and first released in 2008.</dd>

            --fcc-editable-region--
            
            --fcc-editable-region--
        </dl>
    </body> 
</html>
```

