---
id: 6716249b5405164036fd0b0d
title: Costruisci un visualizzatore di ordinamento
challengeType: 25
dashedName: build-a-sorting-visualizer
demoType: onClick
---

# --description--

L'algoritmo Bubble Sort ordina una sequenza di numeri interi confrontando coppie di elementi adiacenti a partire dall'inizio della sequenza. Se il primo elemento è maggiore del secondo, li scambia. Poi procede con la coppia successiva. Quando si raggiunge l'ultimo elemento della sequenza, inizia un nuovo ciclo dall'inizio della sequenza e ripete il processo finché gli elementi non sono ordinati. L'algoritmo si ferma dopo che un ciclo si completa senza scambi.

Per questo laboratorio, ti sono stati forniti tutto l'HTML e il CSS. Userai JavaScript per completare il visualizzatore di Bubble Sort in modo che visualizzi ogni passaggio necessario all'algoritmo Bubble Sort per ordinare un array di cinque numeri interi.

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `generateElement` che restituisce un numero intero casuale tra `1` e `100`, inclusi.
1. Dovresti avere una funzione chiamata `generateArray` che usa la funzione `generateElement` per restituire un array contenente cinque numeri interi casuali.
1. Dovresti avere una funzione chiamata `generateContainer` che crea e restituisce un elemento `div` vuoto.
1. Dovresti avere una funzione chiamata `fillArrContainer` che prende un elemento HTML come primo argomento e un array come secondo argomento.
1. La funzione `fillArrContainer` dovrebbe prendere un elemento come primo parametro e un array di numeri interi come secondo parametro, quindi popolare l'elemento con cinque elementi `span`, ognuno dei quali mostra uno dei numeri dell'array.
1. Dovresti avere una funzione chiamata `isOrdered` che prende due numeri interi e restituisce un booleano che indica se il primo numero è minore o uguale al secondo.
1. Dovresti avere una funzione chiamata `swapElements` che prende un array di numeri interi e un indice numerico.
1. La funzione `swapElements` dovrebbe modificare l'array in loco scambiando l'elemento all'indice passato con quello successivo se `isOrdered` restituisce `false`.
1. Dovresti avere una funzione chiamata `highlightCurrentEls` che prende un elemento HTML e un indice numerico.
1. La funzione `highlightCurrentEls` dovrebbe impostare il `border` del figlio dell'elemento dato all'indice specificato e del figlio immediatamente successivo, assegnando uno stile `dashed`, un colore `red` e una larghezza a tua scelta.
1. Quando clicchi `#generate-btn` dovresti usare la funzione `fillArrContainer` per riempire `#starting-array` con cinque elementi `span`, ognuno con un numero casuale come testo. Se presenti, altri elementi dovrebbero essere rimossi da `#array-container`.
1. Dovresti implementare l'algoritmo Bubble Sort in modo che dopo aver cliccato `#sort-btn`, `#array-container` contenga un elemento `div` per ciascuno dei passaggi richiesti dall'algoritmo Bubble Sort per ordinare l'array iniziale, inclusi `div` che rappresenta l'array iniziale e `div` che rappresenta l'array ordinato. Le funzioni che hai creato finora possono essere utili qui.
1. Ogni `div` dovrebbe contenere cinque elementi `span`, che rappresentano l'array nel suo stato attuale di ordinamento.
1. Dopo aver cliccato `#sort-btn`, `#starting-array` dovrebbe rappresentare il passaggio iniziale con l'array iniziale e i primi due numeri evidenziati.
1. Per ogni passaggio di ordinamento, dovresti usare `highlightCurrentEls` per evidenziare i due numeri che vengono confrontati e scambiarli nel passaggio successivo usando `swapElements`.

# --hints--

Dovresti avere una funzione chiamata `generateElement`.

```js
assert.isFunction(generateElement);
```

La tua funzione `generateElement` dovrebbe restituire un numero intero casuale tra `1` e `100` inclusi.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    assert.strictEqual(generateElement(), 26);
    assert.strictEqual(generateElement(), 9);
} finally {
    randomMocker.restore();
}
```

Dovresti avere una funzione chiamata `generateArray`.

```js
assert.isFunction(generateArray)
```

La tua funzione `generateArray` dovrebbe usare la funzione `generateElement`.

```js
let flag = false;
const temp = generateElement;
generateElement = () => flag = true;
try {
    generateArray();
    assert.isTrue(flag);
} finally {
    generateElement = temp;
}
```

La tua funzione `generateArray` dovrebbe restituire un array contenente cinque numeri interi casuali tra `1` e `100`.

```js
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {
    assert.deepEqual(generateArray(), [26, 9, 58, 23, 38])
} finally {
    randomMocker.restore();
}
```

Dovresti avere una funzione chiamata `generateContainer`.

```js
assert.isFunction(generateContainer);
```

La tua funzione `generateContainer` dovrebbe restituire un elemento `div` vuoto.

```js
const div = generateContainer();
assert.equal(div.tagName, "DIV");
assert.isEmpty(div.children);
```

Dovresti avere una funzione chiamata `fillArrContainer`.

```js
assert.isFunction(fillArrContainer);
```

La tua funzione `fillArrContainer()` dovrebbe prendere un elemento come primo parametro e un array di numeri interi come secondo parametro, quindi popolare l'elemento con cinque elementi `span`, ognuno dei quali mostra uno dei numeri dell'array.

```js
const testDiv = document.createElement("div");
const testArr = [15, 98, 17, 5, 63]
fillArrContainer(testDiv, testArr);
const children = testDiv.children;
assert.lengthOf(children, 5);
Array.from(children).forEach((el, i) => {
    assert.equal(el.tagName, "SPAN");
    assert.equal(el.innerText.trim(), testArr[i])
})
```

Dovresti avere una funzione chiamata `isOrdered`.

```js
assert.isFunction(isOrdered);
```

La tua funzione `isOrdered` dovrebbe prendere due numeri interi e restituire un booleano che indica se il primo numero è minore o uguale al secondo.

```js
assert.lengthOf(isOrdered, 2);
assert.isTrue(isOrdered(2, 60));
assert.isFalse(isOrdered(10, 3));
assert.isTrue(isOrdered(5, 5));
```

Dovresti avere una funzione chiamata `swapElements`.

```js
assert.isFunction(swapElements);
```

La tua funzione `swapElements` dovrebbe prendere un array di numeri interi e un indice numerico come argomenti. Dovrebbe modificare l'array passato in loco scambiando l'elemento all'indice dato con quello successivo se il primo elemento è maggiore del secondo.

```js
const testArr = [22, 4, 87, 47, 33];
swapElements(testArr, 0);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 1);
assert.deepEqual(testArr, [4, 22, 87, 47, 33]);
swapElements(testArr, 2);
assert.deepEqual(testArr, [4, 22, 47, 87, 33]);
swapElements(testArr, 3);
assert.deepEqual(testArr, [4, 22, 47, 33, 87]);
```

Dovresti avere una funzione chiamata `highlightCurrentEls`.

```js
assert.isFunction(highlightCurrentEls);
```

La tua funzione `highlightCurrentEls` dovrebbe dare ai discendenti dell'elemento specificato, situati all'indice dato e all'indice successivo, un bordo che sia `dashed`, `red` e impostato a una larghezza a tua scelta.

```js
const testDiv = document.createElement("div");
document.querySelector("body").appendChild(testDiv)
for (let i = 0; i < 5; i++) {
    testDiv.appendChild(document.createElement("span"));
}
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|#F00|hsl\(0,\s*100%,\s*50%\))/;
const revertBorder = () => {
    for (const el of children) {
        el.style.border = "revert";
    }
}
const children = testDiv.children;

for (let i = 0; i < 3; i++) {
    highlightCurrentEls(testDiv, i);
    for (let j = 0; j < 5; j++) {
        let b = getComputedStyle(children[j]).border;
        if (j == i || j == i + 1) {            
            assert.match(b, redBorderRegex);
            assert.isAbove(parseFloat(b), 0);
        } else {
            assert.notMatch(b, redBorderRegex);
        }
    }
    revertBorder();
}
testDiv.remove();
```

Quando clicchi `#generate-btn` dovresti riempire `#starting-array` con cinque elementi `span`, ognuno con un numero casuale tra `1` e `100` come testo.

```js
const genBtn = document.querySelector("#generate-btn");
genBtn.click();
const children = document.querySelector("#starting-array").querySelectorAll("span");
assert.lengthOf(children, 5);
Array.from(children).forEach(el => {
    assert.equal(el.tagName, "SPAN");
    const num = Number(el.innerText.trim());
    assert.isAtMost(num, 100);
    assert.isAtLeast(num, 1);
})
```

Quando `#starting-array` contiene già un array generato, o `#array-container` contiene l'array ordinato, cliccare `#generate-btn` dovrebbe rimuovere altri elementi in `#array-container`, lasciando solo `#starting-array` con numeri appena generati.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
genBtn.dispatchEvent(new Event("click"));

const prevNumbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));

sortBtn.dispatchEvent(new Event("click"));
genBtn.dispatchEvent(new Event("click"));

const container = document.querySelector("#array-container");
assert.lengthOf(container.children, 1);
const numbers = Array.from(document.querySelector("#starting-array").querySelectorAll("span")).map(el => Number(el.innerText.trim()));
assert.lengthOf(numbers, 5);
assert.isTrue(prevNumbers.some((num, index) => num !== numbers[index]))
```

Dopo aver cliccato `#sort-btn`, `#array-container` dovrebbe contenere tanti elementi `div` quanti sono i passaggi richiesti dall'algoritmo Bubble Sort per ordinare l'array iniziale, inclusi `div` che rappresenta l'array iniziale e `div` che rappresenta l'array ordinato.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
// using randomMocker to be sure that the starting array requires 13 steps to be sorted
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.lengthOf(container.children, 13)
    Array.from(container.children).forEach(el => {assert.equal(el.tagName, "DIV")})
} finally {
    randomMocker.restore();
}
```

Dopo aver cliccato `#sort-btn`, ogni `div` dentro `#array-container` dovrebbe contenere cinque `span`, ognuno con un numero come testo, e disposti per rappresentare i passaggi richiesti dall'algoritmo Bubble Sort per ordinare l'array iniziale.

```js
const finalArr = [9, 23, 26, 38, 58];
const arrays = [
    [26, 9, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 58, 23, 38],
    [9, 26, 23, 58, 38],
    [9, 26, 23, 38, 58],
    [9, 26, 23, 38, 58],
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr,
    finalArr
]
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
const randomMocker = new __helpers.RandomMocker();
randomMocker.mock();
try {    
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    const container = document.querySelector("#array-container");
    assert.isNotEmpty(container.children);
    Array.from(container.children).forEach((el, i) => {
        Array.from(el.children).forEach((j, k) => {
            assert.strictEqual(Number(j.innerText.trim()), arrays[i][k])
        })
    })
} finally {
    randomMocker.restore();
}
```

Quando clicchi `#sort-btn`, dovresti usare la funzione `highlightCurrentEls` per evidenziare gli elementi confrontati in ogni passaggio.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");
let flag = false;
const temp = highlightCurrentEls;
highlightCurrentEls = () => flag = true;
try {
    genBtn.dispatchEvent(new Event("click"));
    sortBtn.dispatchEvent(new Event("click"));
    assert.isTrue(flag);
} finally {
    highlightCurrentEls = temp;
}
```

Dopo aver cliccato `#sort-btn`, `#starting-array` dovrebbe rappresentare il passaggio iniziale con l'array iniziale e i primi due numeri evidenziati usando `highlightCurrentEls`.

```js
const genBtn = document.querySelector("#generate-btn");
const sortBtn = document.querySelector("#sort-btn");   
genBtn.dispatchEvent(new Event("click"));
sortBtn.dispatchEvent(new Event("click"));
const firstContainer = document.querySelector("#starting-array");
const children = firstContainer.children
const redBorderRegex = /dashed (rgb\(255,\s*0,\s*0\)|#FF0000|hsl\(0,\s*100%,\s*50%\))/;
assert.match(getComputedStyle(children[0]).border, redBorderRegex);
assert.match(getComputedStyle(children[1]).border, redBorderRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <meta http-equiv="X-UA-Compatible" content="ie=edge">
    <title>Sorting Visualizer</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main>
        <div id="array-container">
            <div id="starting-array"></div>
        </div>
        <div id="btn-container">
            <button id="generate-btn" type="button">Generate Array</button>
            <button id="sort-btn" type="button">Sort Array</button>
        </div>
    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
    box-sizing: border-box;
}

main {
    height: 100vh;
    display: flex;
    justify-content: center;
    flex-direction: column;
    align-items: center;
}

#array-container {
    max-height: 95vh;
    display: flex;
    flex-direction: column;
    flex-wrap: wrap;
    gap: 2px;

}

#array-container>div {
    min-width: 8rem;
    height: 2rem;
    box-shadow: rgba(50, 50, 93, 0.25) 0px 2px 5px -1px, rgba(0, 0, 0, 0.3) 0px 1px 3px -1px;
    border-radius: 10px;
    margin-bottom: 0.2rem;
    border: 2px solid darkgray;
    display: flex;
    justify-content: space-evenly;
    align-items: center;
}

#starting-array {
    border: 4px solid darkblue !important;
}

#btn-container {
    display: flex;
    justify-content: space-around;
}

button {
    padding: 2px;
    margin: 5px;
}

#sort-btn {
   display: none 
}

span {
    border-radius: 2px;
    padding: 0.5px;
    margin: 0
}

@media (min-width: 430px) {
  #array-container>div {
    min-width: 12rem;    
  }
  span {
    padding: 1px;
    margin: 1px;
  }
}
```

```js
const arrayContainer = document.getElementById("array-container");
const generateArrayBtn = document.getElementById("generate-btn");
const sortArrayBtn = document.getElementById("sort-btn");
const arrayLength = 5;
const minVal = 1;
const maxVal = 100;
const isStart = () => arrayContainer.children.length === 1;
const clearArrayContainer = () => { arrayContainer.innerHTML = '<div id="starting-array"></div>' };
const showSortBtn = () => {sortArrayBtn.style.display = "inline-block"};
const hideSortBtn = () => {sortArrayBtn.style.display = "none"};
const generateElement = () => Math.floor(Math.random() * maxVal + minVal);
const generateArray = () => Array.from({ length: arrayLength }, generateElement);
const fillArrContainer = (container, arr) => {
    container.innerHTML = "";
    arr.forEach(i => {
        const el = document.createElement("span");
        el.innerText = i;
        el.id = `number-${i}`
        container.appendChild(el);
    })
}
const isOrdered = (el1, el2) => el1 <= el2;

const swapElements = (arr, n = 0) => {
    if (n < arr.length - 1 && !isOrdered(arr[n], arr[n + 1])) {
        const temp = arr[n];
        arr[n] = arr[n + 1];
        arr[n + 1] = temp;
    }
}
const getLastChildren = () => arrayContainer.lastElementChild;
const getLastArr = () => {
    const els = Array.from(getLastChildren().children);
    const arr = els.map(el => Number(el.id.replace("number-", "")));
    return arr;
}
const generateContainer = () => {
    const container = document.createElement("div");
    arrayContainer.appendChild(container);
    return container;
}
const highlightCurrentEls = (container, n = 0) => {
    const children = container.children;
    children[n].style.border = "2px dashed red";
    children[n + 1].style.border = "2px dashed red";
}

const highlightSorted = () => {
    getLastChildren().style.border = "4px solid green";
}

const bubbleSort = () => {
    let swapped = true;
    while (swapped) {
        const startingArr = getLastArr();
        startingArr.forEach((_, i) => {
            if (i + 1 < arrayLength) {
                highlightCurrentEls(getLastChildren(), i);
                const arr = getLastArr();
                swapElements(arr, i);
                fillArrContainer(generateContainer(), arr);
            }
        })
        const lastArr = getLastArr()
        if (startingArr.every((el, i) => el === lastArr[i])) swapped = false;
    }
}


generateArrayBtn.addEventListener("click", () => {
    if (!isStart()) {
        clearArrayContainer();
    }
    fillArrContainer(document.getElementById("starting-array"), generateArray());
    showSortBtn();
});

sortArrayBtn.addEventListener("click", () => {
    bubbleSort();
    highlightSorted();
    hideSortBtn();
})
```
