---
id: 6852db755c9c1274d75aae9a
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 시의 각 줄에 숨겨진 핀 숫자를 추출하는 핀 추출기를 만듭니다.

`pin_extractor`이라는 매개변수를 가진 `poem`라는 함수를 만드는 것부터 시작하세요. 빈 함수는 오류가 발생하므로, 유효하게 만들기 위해 잠시 함수 안에 `pass`를 추가하세요.

# --hints--

`pin_extractor`라는 이름의 함수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('pin_extractor')
    `)
})
```

`pin_extractor` 함수는 `poem`라는 매개변수를 가져야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_args('poem')
    `)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
