---
id: 6852dd8486c7ee0171322950
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

디코딩하는 동안 비밀 코드를 저장할 변수가 필요하니 `pass`를 제거하고 함수 안에 `secret_code`라는 변수를 만들고 빈 문자열을 할당하세요.

# --hints--

`secret_code` 함수 안에 `pin_extractor` 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_code')
    `)
})
```

변수 `secret_code`는 빈 문자열 값을 가져야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_code').is_equivalent('secret_code = ""')
    `)
})
```

코드에 더 이상 `pass`가 없어야 합니다.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').has_pass()
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
