---
id: 68530231c05b93324e071104
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

함수를 작업하는 동안 디버깅 용도로 함수 호출이 있으면 유용하니, 함수 밖에 이 시를 담은 `poem` 변수를 만드세요:

```md
Stars and the moon
shine in the sky
white and bright
until the end of the night
```

여러 줄 문자열을 사용할 수 있습니다.

그다음, `poem` 변수를 인수로 하여 함수를 호출하세요.

# --hints--

`poem`라는 전역 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_variable('poem')
    `)
})
```

`poem` 변수는 `Stars and the moon\nshine in the sky\nwhite and bright\nuntil the end of the night`와 같은 값을 가져야 합니다.

```js
({
    test: () => runPython(`
    pm = """Stars and the moon
    shine in the sky
    white and bright
    until the end of the night"""
    assert _Node(_code).find_variable('poem').is_equivalent(f'poem = """{pm}"""')
    `)
})
```

`pin_extractor` 함수를 `poem`를 인수로 호출해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('pin_extractor(poem)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''

--fcc-editable-region--

--fcc-editable-region--
```
