---
id: 68531fbb934d735415053274
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

함수 호출을 추가했으니 함수 내부에서 발생한 오류를 터미널에서 확인할 수 있고, 항상 `print`를 사용해 값들을 확인할 수 있습니다.

핀 번호의 숫자가 각 줄에 숨겨져 있으니, 함수 내부에서 `split` 메서드를 사용해 문자열을 줄 단위 리스트로 나누세요. 줄은 개행 문자(`\n`)를 기준으로 나누고, 결과 리스트를 `lines`라는 변수에 할당하세요.

# --hints--

`lines` 함수 안에 `pin_extractor`라는 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('lines')
    `)
})
```

변수 `lines`는 `poem.split('\n')` 값을 가져야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('lines').is_equivalent('lines = poem.split("\\\\n")')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)
```
