---
id: 68557bd29f7f3c1880f98836
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

루프의 마지막 줄에서 `print` 호출을 업데이트하여 `words` 리스트의 `line_index` 인덱스에 있는 단어를 출력하세요.

# --hints--

루프 끝에 `print(words[line_index])`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words[line_index])')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words)
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
