---
id: 68559f3561d7df287586142e
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

이제 핀 추출기는 작동하지만, 더 효율적일 수 있습니다! 여러 시에서 동시에 핀을 추출할 수 있습니다!

전역 범위에 시 두 편 `poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'`와 `poem3 = 'There\nonce\nwas\na\ndragon'`를 더 추가하고, 지금은 `print`를 호출하는 `pin_extractor`를 주석 처리하세요.


# --hints--

변수 `poem2`와 `poem3`가 존재하고 올바른 값을 가져야 합니다.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_variable('poem2')
        assert _Node(_code).has_variable('poem3')
        assert _Node(_code).find_variable('poem2').is_equivalent("poem2 = 'The grass is green\\\\nhere and there\\\\nhoping for rain\\\\nbefore it turns yellow'")
        assert _Node(_code).find_variable('poem3').is_equivalent("poem3 = 'There\\\\nonce\\\\nwas\\\\na\\\\ndragon'")
    `)
})
```

`print(pin_extractor(poem))` 호출을 주석 처리해야 합니다.

```js
({
    test: () => runPython(`
        assert not _Node(_code).has_call('print(pin_extractor(poem))')`)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
        else:
            secret_code += '0'
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

--fcc-editable-region--


print(pin_extractor(poem))
--fcc-editable-region--

```
