---
id: 6855918e87c33c6979d21d3f
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이제 함수가 시에서 코드를 잘 추출하고 있지만, 고려해야 할 예외 상황이 있습니다: 시의 한 줄이 예상보다 짧을 수 있는데, 이 경우 함수가 오류를 일으킬 수 있습니다(시에서 `bright` 단어를 제거해 시도해 보세요).

`secret_code += str(len(words[line_index]))` 목록에 충분한 단어가 있는지 확인하는 `if` 문 안에 `words` 줄을 넣으세요.

# --hints--

반복문 안에 `if` 문이 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0]
    `)
})
```

if 문 조건은 `len(words) > line_index` 이어야 합니다.


```js
({
    test: () => runPython(`
    _cond = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[0]
    assert _cond.is_equivalent('len(words) > line_index') or _cond.is_equivalent('line_index < len(words)')
    `)
})
```

`if` 문 안에 `secret_code += str(len(words[line_index]))` 이 있어야 합니다.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_stmt('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
--fcc-editable-region--
        secret_code += str(len(words[line_index]))
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

print(pin_extractor(poem))

```
