---
id: 68553423d775cc60134ac8bb
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

각 줄을 독립적으로 처리해야 합니다: `lines`를 순회하는 루프를 만들고 `line`를 루프 변수로 사용하세요. 루프 안에서 `line`를 출력하세요.

# --hints--

`for` 함수 안에 `pin_extractor` 루프가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0]
    `)
})
```

루프는 `lines`를 순회해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('lines')
    `)
})
```

루프 변수는 `line`여야 합니다 (`for line in lines`).

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line')
    `)
})
```

루프 안에 `print(line)`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
