---
id: 685535bf527b2c77236b3a9a
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

핀의 n번째 숫자는 n번째 줄의 n번째 단어 길이로 숨겨져 있습니다.

n번째 단어의 길이를 찾으려면, 다음 단계로 시의 한 줄을 단어 목록으로 분리하세요: 반복문 안에서 변수 `words`를 만들고 `line`을 `split` 메서드를 사용해 단어로 분리한 값을 할당하세요.

그다음 `print`를 인수로 하는 `words` 호출을 추가하세요.

# --hints--

반복문 안에 `words` 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_variable('words')
    `)
})
```

`words` 변수는 `line.split()` 값을 가져야 합니다.

```js
({
    test: () => runPython(`
    variable = _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_variable('words')
    assert variable.is_equivalent('words = line.split()') or variable.is_equivalent('words = line.split(" ")')
    `)
})
```

`print(words)`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(words)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line in lines:
        print(line)
--fcc-editable-region--
        
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
