---
id: 685578aae65e4106ad2b4300
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

이전 수업에서 배운 것처럼, `enumerate` 함수는 반복 가능한 객체를 순회할 때 인덱스를 추적할 수 있게 해줍니다:

```py
fruits = ['apple', 'plum', 'bananas']

for index, item in enumerate(fruits):
    print(index, item)

# 0 apple
# 1 plum
# 2 bananas
```

비밀 코드의 n번째 숫자는 n번째 줄의 n번째 단어에서 나오므로, 각 줄의 번호가 무엇인지 알아야 합니다.

반복문이 `enumerate(lines)`을 순회하도록 바꾸고, `line` 앞에 `line_index`이라는 또 다른 반복 변수도 추가하세요.

또한 `print` 호출을 `print(line_index, line)`로 업데이트하세요.

# --hints--

반복문은 `enumerate(lines)`을 순회해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_iter().is_equivalent('enumerate(lines)')
    `)
})
```

반복문 변수로 `line_index, line`이 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_for_vars().is_equivalent('line_index, line')
    `)
})
```

반복문 안에 `print(line_index, line)`이 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(line_index, line)')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
--fcc-editable-region--
    for line in lines:
        print(line)
--fcc-editable-region--
        words = line.split()
        print(words)

poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
