---
id: 6859a082915dae1e0eaaf962
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

마지막으로 `secret_codes`를 반환하고, 함수 호출의 주석을 해제한 뒤 인수를 `[poem, poem2, poem3]`로 업데이트하세요. 이제 함수가 여러 시를 한 번에 처리할 수 있고, 여러 개의 핀을 추출할 수 있습니다!

`pin_extractor` 함수가 완성되었습니다.

# --hints--

함수는 마지막에 `secret_codes`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_codes')
    `)
})
```

함수 호출 `print(pin_extractor([poem, poem2, poem3]))`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor([poem, poem2, poem3]))')
    `)
})
```

`pin_extractor([poem, poem2, poem3])`는 `['5202', '3346', '50000']`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert pin_extractor([poem, poem2, poem3]) == ['5202', '3346', '50000']
    `)
})
```

함수는 다른 시 목록에서도 올바르게 작동해야 합니다.

```js
({
    test: () => runPython(`
plum = '''In dusky shade the plum tree grew,
Its branches bent with evening dew.
A single fruit, so ripe and round,
Falls softly to the waiting ground.'''
apple = '''Red or green on branches high,
Like lanterns hung against the sky.
Crisp with autumn's golden breath,
A bite that wakes the world from death.'''
grape = '''Clustered tight on winding vines,
Bathed in sun where silence shines.
Jewel-toned beads in twilight air,
A whispered promise hanging there.'''
assert pin_extractor([plum, apple, grape]) == ['2863', '3885', '9227']
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
    secret_codes = []
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        secret_codes.append(secret_code)
--fcc-editable-region--
        

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))
--fcc-editable-region--

```

# --solutions--

```py
def pin_extractor(poems):
    secret_codes = []
    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        secret_codes.append(secret_code)
    return secret_codes

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

print(pin_extractor([poem, poem2, poem3]))

```
