---
id: 685929ee6d500c6f2bce6751
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

`pin_extractor(poem)`을 `print` 함수에 전달해서 함수가 반환하는 값을 확인하세요.

# --hints--

`print(pin_extractor(poem))` 함수를 호출해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_call('print(pin_extractor(poem))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        secret_code += str(len(words[line_index]))
    return secret_code


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

--fcc-editable-region--
pin_extractor(poem)
--fcc-editable-region--

```
