---
id: 685926db99041c5e619ead41
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

`len` 함수는 문자열의 길이를 알려줍니다. 두 번째 `print` 호출을 업데이트해서 그 단어의 길이를 출력하세요.

# --hints--

반복문의 끝에 `print(len(words[line_index]))`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(words[line_index])
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
