---
id: 6859a037b3d6721c7fbb8055
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

루프 전에 빈 리스트를 시작 값으로 하는 새로운 변수 `secret_codes`를 만드세요.

그다음 `return secret_code` 줄을 `secret_code` 리스트에 `secret_codes`를 추가하는 줄로 바꾸세요.

# --hints--

`secret_codes` 함수 안에 `pin_extractor` 변수가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_variable('secret_codes')
    `)
})
```

변수 `secret_codes`는 빈 리스트 값을 가져야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_variable('secret_codes').is_equivalent('secret_codes = []')
    `)
})
```

루프가 끝날 때 `secret_codes.append(secret_code)`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('secret_codes.append(secret_code)')
    `)
})
```

`secret_code`를 반환하면 안 됩니다.

```js
({
    test: () => runPython(`
    assert not _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_return('secret_code')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poems):
--fcc-editable-region--

    for poem in poems:
        secret_code = ''
        lines = poem.split('\n')
        for line_index, line in enumerate(lines):
            words = line.split()
            if len(words) > line_index:
                secret_code += str(len(words[line_index]))
            else:
                secret_code += '0'
        return secret_code
--fcc-editable-region--

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

poem2 = 'The grass is green\nhere and there\nhoping for rain\nbefore it turns yellow'
poem3 = 'There\nonce\nwas\na\ndragon'

# print(pin_extractor(poem))

```
