---
id: 6859290fc8e07e6b2a8d8523
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

두 개의 `print` 호출을 제거하고 함수 마지막 줄에 `return`를 반환하는 `secret_code` 문을 추가하세요.

# --hints--

`print` 호출이 없어야 합니다.

```js
({
    test: () => runPython(`
    assert not _Node(_code).block_has_call('print')
    `)
})
```

`pin_extractor` 함수는 마지막에 `secret_code`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').has_return('secret_code')
    `)
})
```

`pin_extractor(poem)`는 `5262`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5262'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
--fcc-editable-region--
        print(line_index, line)
        words = line.split()
        print(str(len(words[line_index])))
        secret_code += str(len(words[line_index]))
    
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
