---
id: 68592b42b9c7f27cbe123203
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

시의 세 번째 줄에 세 번째 단어가 빠져 있어서 핀이 예상보다 짧습니다.

단어가 충분하지 않을 때 `else`과 `'0'`를 연결하는 `secret_code` 절을 추가하여 시의 모든 줄이 숫자를 찾는 데 사용되도록 하세요.

# --hints--

`else` 절이 있어야 합니다.


```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_ifs()[0].find_conditions()[1] == _Node()
    `)
})
```

`else` 절에는 `secret_code += '0'`가 포함되어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_if('len(words) > line_index').is_equivalent('if len(words) > line_index:\\n  secret_code += str(len(words[line_index]))\\nelse:\\n  secret_code += \\'0\\'')
    `)
})
```

`pin_extractor(poem)`는 `5202`를 반환해야 합니다.


```js
({
    test: () => runPython(`
    assert pin_extractor(poem) == '5202'
    `)
})
```

`secret_code('There\nonce\nwas\na\ndragon')`는 `50000`를 반환해야 합니다.

```js
({
    test: () => runPython(`
    assert pin_extractor('There\\nonce\\nwas\\na\\ndragon') == '50000'
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        words = line.split()
        if len(words) > line_index:
            secret_code += str(len(words[line_index]))
--fcc-editable-region--
        
--fcc-editable-region--
    return secret_code

poem = """Stars and the moon
shine in the sky
white and
until the end of the night"""

print(pin_extractor(poem))

```
