---
id: 68592730aaf34b61ae279168
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

`len`에서 반환된 길이는 정수입니다. `str` 함수를 사용해 문자열로 변환해야 `secret_code` 문자열에 더할 수 있습니다.

길이를 문자열로 출력하도록 `print` 호출을 업데이트하세요.

# --hints--

반복문의 끝에 `print(str(len(words[line_index])))`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].has_call('print(str(len(words[line_index])))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(len(words[line_index]))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
