---
id: 68592807861cfb658c90f70f
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

이제 `str(len(words[line_index]))`와 `secret_code`을 연결할 수 있습니다. `print` 호출을 제거하고 연결을 위해 증강 할당 `+=`을 사용하세요.

# --hints--

반복문 안에는 `secret_code += str(len(words[line_index]))`가 있어야 합니다.

```js
({
    test: () => runPython(`
    assert _Node(_code).find_function('pin_extractor').find_for_loops()[0].find_bodies()[0].find_aug_variable('secret_code').is_equivalent('secret_code += str(len(words[line_index]))')
    `)
})
```

# --seed--

## --seed-contents--

```py
def pin_extractor(poem):
    secret_code = ''
    lines = poem.split('\n')
    for line_index, line in enumerate(lines):
        print(line_index, line)
        words = line.split()
--fcc-editable-region--
        print(str(len(words[line_index])))
--fcc-editable-region--


poem = """Stars and the moon
shine in the sky
white and bright
until the end of the night"""

pin_extractor(poem)

```
