---
id: 616971b3cd990186b66c99a1
title: Passaggio 8
challengeType: 0
dashedName: step-8
---

# --description--

Ora che il tuo file CSS esterno è pronto, puoi iniziare a stilizzare la pagina.

Come promemoria, ecco come selezionare un elemento paragrafo e allinearlo a destra:

```css
p {
  text-align: right;
}
```

Crea una nuova regola CSS che selezioni l'elemento `h1` e imposta la sua proprietà `text-align` su `center`.

# --hints--

Dovresti usare un selettore `h1`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1'));
```

La tua regola CSS `h1` dovrebbe avere una proprietà `text-align` impostata su `center`.

```js
assert(new __helpers.CSSHelp(document).getStyle('h1')?.textAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
