---
id: 61695d1fbc003856628bf561
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

L'attributo `charset` specifica la codifica dei caratteri usata dal documento. `utf-8` (Unicode Transformation Format – 8-bit) è uno standard di codifica dei caratteri usato per la comunicazione elettronica.

All'interno dell'elemento `head`, annida un elemento `meta` con l'attributo `charset` impostato su `"utf-8"`.

# --hints--

Dovresti avere un elemento `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 1);
```

Il tuo elemento `meta` dovrebbe essere un elemento vuoto, non dovrebbe avere un tag di chiusura `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Il tuo elemento `meta` dovrebbe avere un attributo `charset` impostato sul valore `"utf-8"`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

Il tuo elemento `meta` dovrebbe essere all'interno dell'elemento `head`.

```js
const metaElementRegex = /<head\s*>(?:.|\r|\n)*?<meta(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, metaElementRegex);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
