---
id: 61695c4aad56f95497c19583
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Ricorda che l'elemento `title` fornisce ai motori di ricerca informazioni aggiuntive sulla pagina. Mostra anche il contenuto di quell'elemento `title` in due modi ulteriori:

* nella barra del titolo quando la pagina è aperta  
* nella scheda del browser per la pagina quando ci passi sopra con il mouse. Anche se quella scheda non è attiva, una volta che ci passi sopra, viene mostrato il testo `title`.

All'interno dell'elemento `head`, annida un elemento `title` con il testo `Colored Markers`.

# --hints--

Dovresti avere un tag di apertura `title`.

```js
assert(code.match(/<title\s*>/i));
```

Dovresti avere un tag di chiusura `title`.

```js
assert(code.match(/<\/title\s*>/i));
```

Il tuo progetto dovrebbe avere il titolo `Colored Markers`.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim()?.toLowerCase(), 'colored markers')
```

Ricorda, maiuscole e ortografia sono importanti per il testo del titolo.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Colored Markers');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
