---
id: 616965351e74d4689eb6de30
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Puoi avere più elementi `meta` in una pagina web. Ogni elemento `meta` aggiunge informazioni sulla pagina che non possono essere espresse da altri elementi HTML.

Aggiungi un altro elemento `meta` all'interno del `head`. Dagli un attributo `name` impostato su `"viewport"` e un attributo `content` impostato su `"width=device-width, initial-scale=1.0"` in modo che la tua pagina appaia uguale su tutti i dispositivi.

# --hints--

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert.strictEqual(meta?.length, 2);
```

Il tuo nuovo elemento `meta` dovrebbe essere un elemento void, non dovrebbe avere un tag di chiusura `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```


Il tuo nuovo elemento `meta` dovrebbe avere un attributo `name` impostato su `"viewport"`, e un attributo `content` impostato su `"width=device-width, initial-scale=1.0"`.

```js
const meta = [...document.querySelectorAll('meta')];
const contentAttrRegex = /^\s*width\s*=\s*device-width\s*,\s*initial-scale\s*=\s*1(?:\.0)?\s*$/i;
const target = meta?.find(m => m?.getAttribute('name')?.toLowerCase()?.replace(/\s*/g, '') === 'viewport' && contentAttrRegex.test(m?.getAttribute('content')));
assert.exists(target);
```

I tuoi due elementi `meta` dovrebbero essere dentro l'elemento `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('meta').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8">
    <title>Colored Markers</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
