---
id: 616d3a67ccf800ad94ec89ae
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Ora aggiungerai alcuni elementi che alla fine definirai nello stile come marcatori colorati.

Per prima cosa, all'interno dell'elemento `body`, aggiungi un elemento `div` e imposta il suo attributo `class` su `container`. Assicurati che l'elemento `div` sia sotto l'elemento `h1`.

# --hints--

Dovresti avere un tag di apertura `div`.

```js
assert(code.match(/<div.*>/i));
```

Dovresti avere un tag di chiusura `div`.

```js
assert(code.match(/<\/div\s*>/i));
```

Il tuo elemento `div` dovrebbe essere all'interno del tuo elemento `body`.

```js
assert(document.querySelector('div')?.parentElement?.localName === 'body');
```

Il tuo elemento `div` dovrebbe essere dopo l'elemento `h1`.

```js
assert.exists(document.querySelector('h1 + div'));
```

Il tuo elemento `div` dovrebbe avere un attributo `class` impostato su `container`.

```js
assert(document.querySelector('div')?.classList?.contains('container'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
--fcc-editable-region--
  <body>
    <h1>CSS Color Markers</h1>
  </body>
--fcc-editable-region--
</html>
```

```css
h1 {
  text-align: center;
}
```
