---
id: 616d4a84b756d9c4b8255093
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

È il momento di aggiungere un po' di colore al marker. Ricorda che un modo per aggiungere colore a un elemento è usare una <dfn>parola chiave di colore</dfn> come `black`, `cyan` o `yellow`.

Come promemoria, ecco come selezionare la classe `freecodecamp`:

```css
.freecodecamp {
  
}
```

Crea una nuova regola CSS che selezioni la classe `marker` e imposta la sua proprietà `background-color` su `red`.

**Nota:** Non vedrai alcun cambiamento dopo aver aggiunto il CSS.

# --hints--

Dovresti creare un selettore di classe per selezionare la classe `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

La tua regola CSS `.marker` dovrebbe avere una proprietà `background-color` impostata su `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
