---
id: 616d525007b8c5d8b3308b1c
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Il tuo marcatore sarebbe più bello se fosse centrato nella pagina. Un modo semplice per farlo è usare la `margin` <dfn>proprietà abbreviata</dfn>.

Nel progetto precedente, hai impostato separatamente l’area del margine degli elementi con proprietà come `margin-top` e `margin-left`. La `margin` proprietà abbreviata rende facile impostare più aree di margine contemporaneamente.

Per centrare il marcatore nella pagina, imposta la sua proprietà `margin` su `auto`. Questo imposta `margin-top`, `margin-right`, `margin-bottom` e `margin-left` tutte su `auto`.

# --hints--

La tua regola CSS `.marker` dovrebbe avere una proprietà `margin` impostata su `auto`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.margin === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--
.marker {
  width: 200px;
  height: 25px;
  background-color: red;
}
--fcc-editable-region--

```
