---
id: 61764c602bee6974e7790f35
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Per dare ai marker colori diversi, devi aggiungere una classe unica a ciascuno. Puoi aggiungere classi multiple a un elemento elencandole nell'attributo `class` e separandole con uno spazio. Per esempio, il seguente aggiunge sia la classe `animal` sia la classe `dog` a un elemento `div`.

```html
<div class="animal dog">
```

Se aggiungi classi multiple a un elemento HTML, gli stili delle prime classi elencate potrebbero essere sovrascritti da quelli delle classi successive.

Per cominciare, aggiungi la classe `one` al primo elemento marker `div`.


# --hints--

Dovresti aggiungere la classe `one` al primo elemento marker `div`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('one'));
```

Il tuo primo marker `div` dovrebbe avere le classi `marker` e `one`.

```js
const containerFirstChild = [...document.querySelector('.container')?.children][0];
assert(containerFirstChild?.classList?.contains('marker') && containerFirstChild?.classList?.contains('one'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
--fcc-editable-region--
      <div class="marker">
      </div>
      <div class="marker">
      </div>
      <div class="marker">
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  background-color: red;
  margin: 10px auto;
}

```
