---
id: 6176519636a76b810ab1219a
title: Passaggio 20
challengeType: 0
dashedName: step-20
---

# --description--

Crea una regola CSS che abbia come obiettivo la classe `two` e imposta la sua proprietà `background-color` su `green`.

Inoltre, crea una regola CSS separata che abbia come obiettivo la classe `three` e imposta la sua `background-color` su `blue`.

# --hints--

Dovresti usare un selettore di classe per avere come obiettivo la classe `two`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two'));
```

La tua regola CSS `.two` dovrebbe avere una proprietà `background-color` impostata su `green`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.two')?.backgroundColor === 'green');
```

Dovresti usare un selettore di classe per avere come obiettivo la classe `three`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three'));
```

La tua regola CSS `.three` dovrebbe avere una proprietà `background-color` impostata su `blue`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.three')?.backgroundColor === 'blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

.one {
  background-color: red;
}

--fcc-editable-region--

--fcc-editable-region--

```
