---
id: 617b8b38f32bf2080a140675
title: Passaggio 22
challengeType: 0
dashedName: step-22
---

# --description--

Una funzione è un pezzo di codice che può prendere un input ed eseguire un'azione specifica. La funzione CSS `rgb` accetta valori, o <dfn>argomenti</dfn>, per rosso, verde e blu, e produce un colore:

```css
rgb(red, green, blue);
```

Ogni valore di rosso, verde e blu è un numero da `0` a `255`. `0` significa che c'è lo 0% di quel colore, ed è nero. `255` significa che c'è il 100% di quel colore.

Nella regola CSS `.one`, sostituisci la parola chiave di colore `red` con la funzione `rgb`. Per la funzione `rgb`, imposta il valore per il rosso a `255`, il valore per il verde a `0` e il valore per il blu a `0`.

# --hints--

La tua regola CSS `.one` non dovrebbe usare la parola chiave di colore `red` per impostare la proprietà `background-color`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor !== 'red');
```

La tua regola CSS `.one` dovrebbe avere una proprietà `background-color` impostata su `rgb(255, 0, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(0, 0, 0);
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: red;
}
--fcc-editable-region--

.two {
  background-color: green;
}

.three {
  background-color: blue;
}

```
