---
id: 617bb5624a75e86463b7e638
title: Passaggio 28
challengeType: 0
dashedName: step-28
---

# --description--

<dfn>I colori secondari</dfn> sono i colori che ottieni combinando i colori primari. Potresti aver notato alcuni colori secondari nell'ultimo passaggio mentre cambiavi i valori di rosso, verde e blu.

Per creare il primo colore secondario, il giallo, aggiorna la funzione `rgb` nella regola CSS `.one` per combinare il rosso puro e il verde puro.

# --hints--

La tua regola CSS `.one` dovrebbe avere una proprietà `background-color` impostata su `rgb(255, 255, 0)`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.one')?.backgroundColor === 'rgb(255, 255, 0)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker one">
      </div>
      <div class="marker two">
      </div>
      <div class="marker three">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.one {
  background-color: rgb(255, 0, 0);
}
--fcc-editable-region--

.two {
  background-color: rgb(0, 255, 0);
}

.three {
  background-color: rgb(0, 0, 255);
}

```
