---
id: 61a493ead935c148d2b76312
title: Passaggio 50
challengeType: 0
dashedName: step-50
---

# --description--

La funzione `linear-gradient` è molto flessibile — ecco la sintassi base che userai in questo tutorial:

```css
linear-gradient(gradientDirection, color1, color2, ...);
```

`gradientDirection` è la direzione della linea usata per la transizione. `color1` e `color2` sono argomenti di colore, cioè i colori che verranno usati nella transizione stessa. Possono essere qualsiasi tipo di colore, incluse parole chiave di colore, esadecimali, `rgb` o `hsl`.

Ora applicherai un gradiente dal rosso al verde lungo una linea di 90 gradi al primo marker.

Per prima cosa, nella regola CSS `.red`, imposta la proprietà `background` su `linear-gradient()` e passale il valore `90deg` come `gradientDirection`.

# --hints--

La tua regola CSS `.red` dovrebbe avere una proprietà `background` con il valore `linear-gradient(90deg)`.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.red\{.*?background:linear-gradient\(90deg\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: rgb(255, 0, 0);
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
