---
id: 61a49d15bdbb5e57cc6fd280
title: Passaggio 54
challengeType: 0
dashedName: step-54
---

# --description--

I punti di colore ti permettono di regolare con precisione dove i colori vengono posizionati lungo la linea di gradiente. Sono un’unità di lunghezza come `px` o percentuali che seguono un colore nella funzione `linear-gradient`.

Per esempio, in questo gradiente rosso-nero, la transizione dal rosso al nero avviene al 90% della linea di gradiente, quindi il rosso occupa la maggior parte dello spazio disponibile:

```css
linear-gradient(90deg, red 90%, black);
```

Nella funzione `linear-gradient`, aggiungi un punto di colore `75%` dopo il primo argomento colore rosso. Non aggiungere punti di colore agli altri argomenti colore.

# --hints--

La tua regola CSS `.red` dovrebbe avere una proprietà `background` impostata su `linear-gradient(90deg, rgb(255, 0, 0) 75%, rgb(0, 255, 0), rgb(0, 0, 255))`.

```js
assert.include(['linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))', 'rgba(0,0,0,0)linear-gradient(90deg,rgb(255,0,0)75%,rgb(0,255,0),rgb(0,0,255))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(90deg, rgb(255, 0, 0), rgb(0, 255, 0), rgb(0, 0, 255));
}
--fcc-editable-region--

.green {
  background-color: #007F00;
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
