---
id: 61a5bfe091060f1d6a629dd0
title: Passaggio 64
challengeType: 0
dashedName: step-64
---

# --description--

Anche senza i punti di colore, potresti aver notato che i colori per il marcatore verde cambiano negli stessi punti del marcatore rosso. Il primo colore è all'inizio (0%), il secondo è a metà (50%) e l'ultimo è alla fine (100%) della linea di gradiente.

La funzione `linear-gradient` calcola automaticamente questi valori per te e posiziona i colori in modo uniforme lungo la linea di gradiente per impostazione predefinita.

Nella regola CSS `.red`, rimuovi le tre fermate di colore dalla funzione `linear-gradient` per rendere il codice un po' più pulito.

# --hints--

La tua regola CSS `.red` dovrebbe avere una proprietà `background` impostata su `linear-gradient(180deg, rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27))`.

```js
assert.include(['linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))', 'rgba(0,0,0,0)linear-gradient(rgb(122,74,14),rgb(245,62,113),rgb(162,27,27))repeatscroll0%0%'], new __helpers.CSSHelp(document).getStyle('.red')?.getPropVal('background', true));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.red {
  background: linear-gradient(180deg, rgb(122, 74, 14) 0%, rgb(245, 62, 113) 50%, rgb(162, 27, 27) 100%);
}
--fcc-editable-region--

.green {
  background: linear-gradient(180deg, #55680D, #71F53E, #116C31);
}

.blue {
  background-color: hsl(240, 100%, 50%);
}

```
