---
id: 61b093179e7fc020b43b1bb3
title: Passaggio 73
challengeType: 0
dashedName: step-73
---

# --description--

<dfn>Opacità</dfn> descrive quanto qualcosa sia opaco, cioè non trasparente. Per esempio, un muro solido è opaco e nessuna luce può passare attraverso. Ma un bicchiere è molto più trasparente, e puoi vedere attraverso il vetro dall'altra parte.

Con la proprietà CSS `opacity` puoi controllare quanto un elemento sia opaco o trasparente. Con il valore `0`, o 0%, l'elemento sarà completamente trasparente, mentre con `1.0`, o 100%, l'elemento sarà completamente opaco come avviene per impostazione predefinita.

Nella regola CSS `.sleeve` imposta la proprietà `opacity` su `0.5`.

# --hints--

La tua regola CSS `.sleeve` dovrebbe avere una proprietà `opacity` impostata su `0.5`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.sleeve')?.opacity === '0.5');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker red">
        <div class="sleeve"></div>
      </div>
      <div class="marker green">
      </div>
      <div class="marker blue">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

.container {
  background-color: rgb(255, 255, 255);
  padding: 10px 0;
}

.marker {
  width: 200px;
  height: 25px;
  margin: 10px auto;
}

--fcc-editable-region--
.sleeve {
  width: 110px;
  height: 25px;
  background-color: white;
}
--fcc-editable-region--

.red {
  background: linear-gradient(rgb(122, 74, 14), rgb(245, 62, 113), rgb(162, 27, 27));
}

.green {
  background: linear-gradient(#55680D, #71F53E, #116C31);
}

.blue {
  background: linear-gradient(hsl(186, 76%, 16%), hsl(223, 90%, 60%), hsl(240, 56%, 42%));
}

```
