---
id: 66be24cb4144f955b6bcc550
title: 建置一頁撲克牌
challengeType: 25
dashedName: build-a-page-of-playing-cards
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該建置一個至少顯示三張撲克牌的網頁。
1. 你應該有一個 ID 為 `playing-cards` 的 `main` 元素。
1. 在你的 `#playing-cards` 元素中，你應該至少有三個 `div` 元素，每個元素的類別為 `card`。
1. 在每個 `.card` 元素中，你應該有三個 `div` 元素，第一個的類別為 `left`，第二個的類別為 `middle`，第三個的類別為 `right`。
1. 你的 `#playing-cards` 元素應該使用 flexbox 來水平置中其子元素，允許它們換行，並在它們之間放置 `20px` 的間距。
1. 你的每個 `.card` 元素都應該使用 flexbox 以 `space-between` 來對齊其子元素。
1. 你的每個 `.left` 項目都應該使用 flexbox 項目屬性，使其在父容器的起始位置對齊。
1. 你的每個 `.middle` 項目都應該使用 flexbox 項目屬性，使其在父容器中置中對齊。
1. 你的每個 `.right` 項目都應該使用 flexbox 項目屬性，使其在父容器的末端對齊。
1. 你的每個 `.middle` 元素都應該使用 flexbox 以行方式顯示其子元素。

以下是一些你可以複製並貼上以用於你的卡片的字元：`♠`、`♣`、`♥` 和 `♦`。

# --hints--

你應該有一個 `main` 元素，其 `id` 為 `playing-cards`。

```js
assert.exists(document.querySelector('main#playing-cards'));
```

你應該在你的 `#playing-cards` 元素中至少有三個帶有 `card` 類別的 `div` 元素。

```js
assert.isAtLeast(document.querySelectorAll('#playing-cards > div.card')?.length, 3);
```

你的每個 `.card` 元素都應該有寬度和高度。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  assert.isAbove(el.clientWidth, 0);
  assert.isAbove(el.clientHeight, 0);
});
```

你的每個 `.card` 元素應該有且只有三個 `div` 元素作為子元素。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => {
  const children = el.children;
  assert.lengthOf(children, 3);
  Array.from(children).forEach(child => {
    assert.equal(child.tagName, 'DIV');
  });
});
```

你應該在每個 `.card` 元素中有一個帶有 `left` 類別的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.left')));
```

你應該在每個 `.card` 元素中有一個帶有 `middle` 類別的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.middle')));
```

你應該在每個 `.card` 元素中有一個帶有 `right` 類別的 `div`。

```js
const els = document.querySelectorAll('.card');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.exists(el.querySelector('div.right')));
```

你應該有一個 `#playing-cards` 選擇器，將其元素的 `display` 設為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.display, 'flex');
```

你應該有一個 `#playing-cards` 選擇器，將其元素的 `justify-content` 設定為 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.justifyContent, 'center');
```

你應該有一個 `#playing-cards` 選擇器，將其元素的 `flex-wrap` 設為 `wrap`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.flexWrap, 'wrap');
```

你應該有一個 `#playing-cards` 選擇器，將其元素的 `gap` 設為 `20px`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#playing-cards')?.gap, '20px');
```

你應該有一個 `.card` 選擇器，將其元素的 `display` 設定為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.display, 'flex');
```

你應該有一個 `.card` 選擇器，將其元素的 `justify-content` 設為 `space-between`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.card')?.justifyContent, 'space-between');
```

你應該有一個 `.left` 選擇器，將其元素的 `align-self` 設定為 `flex-start`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.left')?.alignSelf, 'flex-start');
```

你應該有一個 `.middle` 選擇器，將其元素的 `align-self` 設為 `center`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.alignSelf, 'center');
```

你應該有一個 `.right` 選擇器，將其元素的 `align-self` 設定為 `flex-end`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.right')?.alignSelf, 'flex-end');
```

你應該有一個 `.middle` 選擇器，將其元素的 `display` 設為 `flex`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.display, 'flex');
```

你應該在你的 `.middle` 類別中將 `flex-direction` 設定為 `column`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.middle')?.flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Playing Cards</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Straight</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="playing-cards">
    <div class="card">
      <div class="left">
        <div>A</div>
        <div>♠</div>
      </div>
      <div class="middle">
        <div>♠</div>
      </div>
      <div class="right">
        <div class="flipped">♠</div>
        <div class="flipped">A</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>2</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div>♣</div>
        <div class="flipped">♣</div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">2</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>3</div>
        <div>♥</div>
      </div>
      <div class="middle">
        <div>♥</div>
        <div>♥</div>
        <div class="flipped">♥</div>
      </div>
      <div class="right">
        <div class="flipped">♥</div>
        <div class="flipped">3</div>
      </div>
    </div>

    <div class="card">
      <div class="left">
        <div>4</div>
        <div>♣</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♣</div>
          <div>♣</div>
        </div>
        <div class="middle-section">
          <div class="flipped">♣</div>
          <div class="flipped">♣</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♣</div>
        <div class="flipped">4</div>
      </div>
    </div>

    <div class="card red">
      <div class="left">
        <div>5</div>
        <div>♦</div>
      </div>
      <div class="middle">
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
        </div>
        <div class="middle-section">
          <div>♦</div>
          <div>♦</div>
        </div>
      </div>
      <div class="right">
        <div class="flipped">♦</div>
        <div class="flipped">5</div>
      </div>
    </div>
  </main>
</body>

</html>
```

```css
body {
  background-color: Thistle;
  min-height: 100vh;
}

.red {
  color: red;
}

.flipped {
  transform: rotate(180deg);
}

#playing-cards {
  margin-top: 100px;
  display: flex;
  justify-content: center;
  gap: 20px;
  flex-wrap: wrap;
}

.card {
  background-color: white;
  border: 1px solid black;
  border-radius: 10px;
  box-shadow: 2px 2px 2px gray;
  min-width: 250px;
  max-width: 250px;
  height: 350px;
  font-size: 30px;
  font-weight: bold;
  padding: 5px;
  box-sizing: border-box;
  display: flex;
  justify-content: space-between;
  text-align: center;
}

.left {
  align-self: flex-start;
}

.middle {
  align-self: center;
  display: flex;
  font-size: 80px;
  flex-direction: column;
}

.right {
  align-self: flex-end;
}

.middle-section {
  display: flex;
  justify-content: center;
}
```
