---
id: 5a858944d96184f06fd60d61
title: Erstellen Sie Ihr erstes CSS-Grid
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cqwREC4'
forumTopicId: 301129
dashedName: create-your-first-css-grid
---

# --description--

Verwandeln Sie jedes HTML-Element in einen Grid-Container, indem Sie seine `display`-Eigenschaft auf `grid` setzen. Dadurch können Sie alle anderen Eigenschaften verwenden, die mit CSS Grid verbunden sind.

**Hinweis:** Im CSS Grid wird das übergeordnete Element als <dfn>Container</dfn> bezeichnet und seine Kinder heißen <dfn>Items</dfn>.

# --instructions--

Ändern Sie die Anzeige des div mit der `container`-Klasse auf `grid`.

# --hints--

Die `container`-Klasse sollte eine `display`-Eigenschaft mit dem Wert `grid` haben.

```js
assert.match(code, /.container\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    display: grid;
  }
</style>
```
