---
id: 5a9036ee38fddaf9a66b5d35
title: Erstellen Sie einen Spaltenabstand mit grid-column-gap
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cVZ8vfD'
forumTopicId: 301124
dashedName: create-a-column-gap-using-grid-column-gap
---

# --description--

Bisher waren die Spalten in den von Ihnen erstellten Grids alle dicht aneinander. Manchmal möchten Sie jedoch einen Abstand zwischen den Spalten haben. Um einen Abstand zwischen den Spalten hinzuzufügen, verwenden Sie die `grid-column-gap`-Eigenschaft so:

```css
grid-column-gap: 10px;
```

Dies erzeugt 10px Leerraum zwischen allen Spalten.

# --instructions--

Geben Sie den Spalten im Grid einen `20px`-Abstand.

# --hints--

Die `container`-Klasse sollte eine `grid-column-gap`-Eigenschaft mit dem Wert `20px` haben.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-column-gap\s*?:\s*?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

   
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-column-gap: 20px;
  }
</style>
```
