---
id: 5a9036e138fddaf9a66b5d33
title: Reihen mit grid-template-rows hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cbp9Pua'
forumTopicId: 301119
dashedName: add-rows-with-grid-template-rows
---

# --description--

Das Grid, das Sie in der letzten Aufgabe erstellt haben, legt die Anzahl der Reihen automatisch fest. Um die Reihen manuell anzupassen, verwenden Sie die `grid-template-rows`-Eigenschaft auf dieselbe Weise, wie Sie `grid-template-columns` in der vorherigen Aufgabe verwendet haben.

# --instructions--

Fügen Sie dem Grid zwei Reihen hinzu, die jeweils `50px` hoch sind.

# --hints--

Die `container`-Klasse sollte eine `grid-template-rows`-Eigenschaft mit zwei Einheiten `50px` haben.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-rows\s*?:\s*?50px\s*?50px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 100px 100px 100px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-rows: 50px 50px;
  }
</style>
```
