---
id: 5a90374338fddaf9a66b5d3a
title: Ein Element horizontal mit justify-self ausrichten
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpKHq'
forumTopicId: 301122
dashedName: align-an-item-horizontally-using-justify-self
---

# --description--

Im CSS Grid befindet sich der Inhalt jedes Elements in einer Box, die als <dfn>Zelle</dfn> bezeichnet wird. Sie können die horizontale Position des Inhalts innerhalb seiner Zelle mit der `justify-self`-Eigenschaft an einem Grid-Element ausrichten. Standardmäßig hat diese Eigenschaft den Wert `stretch`, wodurch der Inhalt die gesamte Breite der Zelle ausfüllt. Diese CSS Grid-Eigenschaft akzeptiert auch andere Werte:

`start`: richtet den Inhalt links in der Zelle aus,

`center`: richtet den Inhalt in der Mitte der Zelle aus,

`end`: richtet den Inhalt rechts in der Zelle aus.

# --instructions--

Verwenden Sie die `justify-self`-Eigenschaft, um das Element mit der Klasse `item2` zu zentrieren.

# --hints--

Die `item2`-Klasse sollte eine `justify-self`-Eigenschaft mit dem Wert `center` haben.

```js
assert.match(
  code,
  /.item2\s*?{[\s\S]*justify-self\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }

  .item2 {
    background: LightSalmon;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item2 {
    justify-self: center;
  }
</style>
```
