---
id: 5a90376038fddaf9a66b5d3c
title: Alle Elemente horizontal mit justify-items ausrichten
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cJbpECn'
forumTopicId: 301120
dashedName: align-all-items-horizontally-using-justify-items
---

# --description--

Manchmal möchten Sie, dass alle Elemente in Ihrem CSS Grid dieselbe Ausrichtung haben. Sie können die zuvor gelernten Eigenschaften verwenden und sie einzeln ausrichten, oder Sie können alle Elemente auf einmal horizontal ausrichten, indem Sie `justify-items` auf Ihren Grid-Container anwenden. Diese Eigenschaft kann alle Werte annehmen, die Sie in den vorherigen beiden Aufgaben gelernt haben, der Unterschied besteht darin, dass sie **alle** Elemente in unserem Grid zur gewünschten Ausrichtung bewegt.

# --instructions--

Verwenden Sie diese Eigenschaft, um alle unsere Elemente horizontal zu zentrieren.

# --hints--

Die `container`-Klasse sollte eine `justify-items`-Eigenschaft mit dem Wert `center` haben.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*justify-items\s*?:\s*?center\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    justify-items: center;
  }
</style>
```
