---
id: 5a9036ee38fddaf9a66b5d37
title: Lücken schneller mit grid-gap hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/ca2qVtv'
forumTopicId: 301118
dashedName: add-gaps-faster-with-grid-gap
---

# --description--

`grid-gap` ist eine Kurzschreibweise-Eigenschaft für `grid-row-gap` und `grid-column-gap` aus den beiden vorherigen Aufgaben, die bequemer zu verwenden ist. Wenn `grid-gap` einen Wert hat, erzeugt sie eine Lücke zwischen allen Reihen und Spalten. Hat sie jedoch zwei Werte, wird der erste Wert verwendet, um die Lücke zwischen den Reihen zu setzen, und der zweite Wert für die Spalten.

# --instructions--

Verwenden Sie `grid-gap`, um eine `10px`-Lücke zwischen den Reihen und eine `20px`-Lücke zwischen den Spalten einzufügen.

# --hints--

Die `container`-Klasse sollte eine `grid-gap`-Eigenschaft haben, die eine `10px`-Lücke zwischen den Reihen und eine `20px`-Lücke zwischen den Spalten einfügt.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-gap\s*?:\s*?10px\s+?20px\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .d1 {
    background: LightSkyBlue;
  }
  .d2 {
    background: LightSalmon;
  }
  .d3 {
    background: PaleTurquoise;
  }
  .d4 {
    background: LightPink;
  }
  .d5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }
</style>
<div class="container">
  <div class="d1">1</div>
  <div class="d2">2</div>
  <div class="d3">3</div>
  <div class="d4">4</div>
  <div class="d5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-gap: 10px 20px;
  }
</style>
```
