---
id: 5a94fe3669fb03452672e45f
title: Wiederholungen mit der Funktion repeat reduzieren
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cQvqyHR'
forumTopicId: 301133
dashedName: reduce-repetition-using-the-repeat-function
---

# --description--

Als Sie `grid-template-columns` und `grid-template-rows` verwendet haben, um die Struktur eines Grids zu definieren, haben Sie für jede erstellte Zeile oder Spalte einen Wert eingegeben.

Angenommen, Sie möchten ein Grid mit 100 Zeilen gleicher Höhe. Es ist nicht sehr praktisch, 100 Werte einzeln einzufügen. Zum Glück gibt es eine bessere Methode – indem Sie die `repeat`-Funktion verwenden, um anzugeben, wie oft Ihre Spalte oder Zeile wiederholt werden soll, gefolgt von einem Komma und dem Wert, den Sie wiederholen möchten.

Hier ist ein Beispiel, das das 100-Zeilen-Grid erstellt, wobei jede Zeile 50px hoch ist.

```css
grid-template-rows: repeat(100, 50px);
```

Sie können auch mehrere Werte mit der repeat-Funktion wiederholen und die Funktion unter anderen Werten einfügen, wenn Sie eine Grid-Struktur definieren. So sieht das aus:

```css
grid-template-columns: repeat(2, 1fr 50px) 20px;
```

Das entspricht:

```css
grid-template-columns: 1fr 50px 1fr 50px 20px;
```

**Hinweis:** Die `1fr 50px` wird zweimal wiederholt, gefolgt von 20px.

# --instructions--

Verwenden Sie `repeat`, um Wiederholungen in der `grid-template-columns`-Eigenschaft zu entfernen.

# --hints--

Die `container`-Klasse sollte eine `grid-template-columns`-Eigenschaft haben, die so eingestellt ist, dass 3 Spalten mit der Breite `1fr` wiederholt werden.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?1fr\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: 1fr 1fr 1fr;

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, 1fr);
  }
</style>
```
