---
id: 5a94fe4469fb03452672e460
title: Größe von Elementen mit der Funktion minmax begrenzen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cD97RTv'
forumTopicId: 301131
dashedName: limit-item-size-using-the-minmax-function
---

# --description--

Es gibt eine weitere eingebaute Funktion, die Sie mit `grid-template-columns` und `grid-template-rows` verwenden können, die `minmax` heißt. Sie dient dazu, die Größe von Elementen zu begrenzen, wenn sich der Grid-Container in der Größe ändert. Dafür müssen Sie den zulässigen Größenbereich für Ihr Element angeben. Hier ist ein Beispiel:

```css
grid-template-columns: 100px minmax(50px, 200px);
```

Im obigen Code ist `grid-template-columns` so eingestellt, dass zwei Spalten erstellt werden; die erste ist 100px breit, und die zweite hat eine Mindestbreite von 50px und eine Maximalbreite von 200px.

# --instructions--

Verwenden Sie die `minmax`-Funktion, um in der `1fr`-Funktion das `repeat` durch eine Spaltengröße zu ersetzen, die eine Mindestbreite von `90px` und eine Maximalbreite von `1fr` hat, und ändern Sie die Größe des Vorschaufensters, um die Wirkung zu sehen.

# --hints--

Die `container`-Klasse sollte eine `grid-template-columns`-Eigenschaft haben, die auf repeat 3 Spalten mit einer Mindestbreite von `90px` und einer Maximalbreite von `1fr` eingestellt ist.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?3\s*?,\s*?minmax\s*?\(\s*?90px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, 1fr);

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .container {
    grid-template-columns: repeat(3, minmax(90px, 1fr));
  }
</style>
```
