---
id: 5a94fe5469fb03452672e461
title: Flexible layouts mit auto-fill erstellen
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cmzdycW'
forumTopicId: 301126
dashedName: create-flexible-layouts-using-auto-fill
---

# --description--

Die repeat-Funktion bietet eine Option namens <dfn>auto-fill</dfn>. Damit können Sie automatisch so viele Reihen oder Spalten Ihrer gewünschten Größe einfügen, wie es die Größe des Containers zulässt. Sie können flexible Layouts erstellen, wenn Sie `auto-fill` mit `minmax` kombinieren, so wie hier:

```css
repeat(auto-fill, minmax(60px, 1fr));
```

Wenn sich die Größe des Containers ändert, fügt diese Einstellung weiterhin 60px-Spalten ein und dehnt sie aus, bis eine weitere Spalte eingefügt werden kann. **Hinweis:** Wenn Ihr Container nicht alle Elemente in einer Reihe aufnehmen kann, werden sie in eine neue Reihe darunter verschoben.

# --instructions--

Verwenden Sie im ersten Grid `auto-fill` mit `repeat`, um das Grid mit Spalten zu füllen, die eine Mindestbreite von `60px` und eine Maximalbreite von `1fr` haben. Ändern Sie dann die Größe der Vorschau, um auto-fill in Aktion zu sehen.

# --hints--

Die `container`-Klasse sollte eine `grid-template-columns`-Eigenschaft mit `repeat` und `auto-fill` haben, die das Grid mit Spalten füllt, die eine Mindestbreite von `60px` und eine Maximalbreite von `1fr` haben.

```js
assert.match(
  code,
  /.container\s*?{[\s\S]*grid-template-columns\s*?:\s*?repeat\s*?\(\s*?auto-fill\s*?,\s*?minmax\s*?\(\s*?60px\s*?,\s*?1fr\s*?\)\s*?\)\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(3, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: LightGray;
    display: grid;
    /* Only change code below this line */

    grid-template-columns: repeat(auto-fill, minmax(60px, 1fr));

    /* Only change code above this line */
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }

  .container2 {
    font-size: 40px;
    min-height: 100px;
    width: 100%;
    background: Silver;
    display: grid;
    grid-template-columns: repeat(3, minmax(60px, 1fr));
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>
<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
<div class="container2">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
