---
id: 5a94fe8569fb03452672e464
title: Raster innerhalb von Rastern erstellen
challengeType: 0
forumTopicId: 301128
dashedName: create-grids-within-grids
---

# --description--

Ein Element in ein Raster zu verwandeln, beeinflusst nur das Verhalten seiner direkten Nachfahren. Wenn Sie also einen direkten Nachfahren in ein Raster verwandeln, erhalten Sie ein Raster innerhalb eines Rasters.

Zum Beispiel erstellen Sie durch das Setzen der `display`- und `grid-template-columns`-Eigenschaften des Elements mit der `item3`-Klasse ein Raster innerhalb Ihres Rasters.

# --instructions--

Verwandeln Sie das Element mit der `item3`-Klasse in ein Raster mit zwei Spalten mit einer Breite von `auto` und `1fr` unter Verwendung von `display` und `grid-template-columns`.

# --hints--

Die `item3`-Klasse sollte eine `grid-template-columns`-Eigenschaft mit den Werten `auto` und `1fr` haben.

```js
assert.match(
  code,
  /.item3\s*?{[\s\S]*grid-template-columns\s*?:\s*?auto\s*?1fr\s*?;[\s\S]*}/gi
);
```

Die `item3`-Klasse sollte eine `display`-Eigenschaft mit dem Wert `grid` haben.

```js
assert.match(code, /.item3\s*?{[\s\S]*display\s*?:\s*?grid\s*?;[\s\S]*}/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  .container {
    font-size: 1.5em;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: auto 1fr;
    grid-template-rows: auto 1fr auto;
    grid-gap: 10px;
    grid-template-areas:
      'advert header'
      'advert content'
      'advert footer';
  }
  .item1 {
    background: LightSkyBlue;
    grid-area: header;
  }

  .item2 {
    background: LightSalmon;
    grid-area: advert;
  }

  .item3 {
    background: PaleTurquoise;
    grid-area: content;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .item4 {
    background: lightpink;
    grid-area: footer;
  }

  .itemOne {
    background: PaleGreen;
  }

  .itemTwo {
    background: BlanchedAlmond;
  }
</style>

<div class="container">
  <div class="item1">header</div>
  <div class="item2">advert</div>
  <div class="item3">
    <div class="itemOne">paragraph1</div>
    <div class="itemTwo">paragraph2</div>
  </div>
  <div class="item4">footer</div>
</div>
```

# --solutions--

```html
<style>
  .item3 {
    grid-template-columns: auto 1fr;
    display: grid;
  }
</style>
```
