---
id: 5a94fe1369fb03452672e45d
title: Elemente in Grid-Bereichen mit der Eigenschaft grid-area platzieren
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/cRrqmtV'
forumTopicId: 301132
dashedName: place-items-in-grid-areas-using-the-grid-area-property
---

# --description--

Nachdem Sie eine Bereichsvorlage für Ihren Grid-Container erstellt haben, wie in der vorherigen Aufgabe gezeigt, können Sie ein Element in Ihrem benutzerdefinierten Bereich platzieren, indem Sie den Namen verwenden, den Sie ihm gegeben haben. Dazu verwenden Sie die `grid-area`-Eigenschaft an einem Element wie folgt:

```css
.item1 {
  grid-area: header;
}
```

Dies teilt dem Grid mit, dass Sie möchten, dass die `item1`-Klasse in dem Bereich mit dem Namen `header` platziert wird. In diesem Fall verwendet das Element die gesamte obere Zeile, weil diese ganze Zeile als `header`-Bereich benannt ist.

# --instructions--

Platzieren Sie ein Element mit der `item5`-Klasse im `footer`-Bereich mithilfe der `grid-area`-Eigenschaft.

# --hints--

Die `item5`-Klasse sollte eine `grid-area`-Eigenschaft haben, deren Wert `footer` ist.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?footer\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */

    
    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: footer;
  }
</style>
```
