---
id: 5a94fe0569fb03452672e45c
title: Das Raster in einen Bereichs-Template aufteilen
challengeType: 0
forumTopicId: 301130
dashedName: divide-the-grid-into-an-area-template
---

# --description--

Sie können Zellen Ihres Rasters zu einem <dfn>Bereich</dfn> zusammenfassen und dem Bereich einen eigenen Namen geben. Das machen Sie, indem Sie `grid-template-areas` auf den Container anwenden, so wie hier:

```css
grid-template-areas:
  'header header header'
  'advert content content'
  'advert footer footer';
```

Der obige Code fasst die Zellen des Rasters in vier Bereiche zusammen: `header`, `advert`, `content` und `footer`. Jedes Wort steht für eine Zelle, und jedes Paar Anführungszeichen steht für eine Zeile.

# --instructions--

Ändern Sie das Template so, dass der Bereich `footer` die gesamte untere Zeile einnimmt. Die Definition der Bereiche hat jetzt noch keine sichtbare Auswirkung. Später werden Sie ein Element so einstellen, dass es einen Bereich verwendet, um zu sehen, wie das funktioniert.

# --hints--

Die `container`-Klasse sollte eine `grid-template-areas`-Eigenschaft haben, die dem Beispiel ähnelt, aber mit dem Bereich `footer`, der die gesamte untere Zeile einnimmt.

```js
assert.match(
  __helpers.removeCssComments(code),
  /.container\s*?{[\s\S]*grid-template-areas\s*?:\s*?["|']\s*?header\s*?header\s*?header\s*?["|']\s*?["|']\s*?advert\s*?content\s*?content\s*?["|']\s*?["|']\s*?footer\s*?footer\s*?footer\s*?["|']\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
    grid-template-areas:
    /* Only change code below this line */
      'header header header'
      'advert content content'
      'advert footer footer';
    /* Only change code above this line */
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }
  .item5 {
    background: PaleGreen;
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;

    grid-template-areas:
      'header header header'
      'advert content content'
      'footer footer footer';
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```
