---
id: 5a94fe2669fb03452672e45e
title: Verwenden von grid-area ohne Erstellen einer Areas-Vorlage
challengeType: 0
videoUrl: 'https://scrimba.com/p/pByETK/c6N7VhK'
forumTopicId: 301135
dashedName: use-grid-area-without-creating-an-areas-template
---

# --description--

Die `grid-area`-Eigenschaft, die Sie in der letzten Aufgabe gelernt haben, kann auf eine andere Weise verwendet werden. Wenn Ihr Grid keine Areas-Vorlage zur Referenz hat, können Sie für ein Element dynamisch einen Bereich erstellen, der so platziert wird:

```css
item1 {
  grid-area: 1/1/2/4;
}
```

Dabei werden die Zeilennummern verwendet, die Sie zuvor gelernt haben, um zu definieren, wo der Bereich für dieses Element liegt. Die Zahlen im obigen Beispiel stehen für diese Werte:

```css
grid-area: horizontal line to start at / vertical line to start at / horizontal
  line to end at / vertical line to end at;
```

Das Element im Beispiel nimmt also die Zeilen zwischen den Linien 1 und 2 sowie die Spalten zwischen den Linien 1 und 4 ein.

# --instructions--

Verwenden Sie die `grid-area`-Eigenschaft, um das Element mit der `item5`-Klasse zwischen der dritten und vierten horizontalen Linie und zwischen der ersten und vierten vertikalen Linie zu platzieren.

# --hints--

Die `item5`-Klasse sollte eine `grid-area`-Eigenschaft haben, damit sie den gesamten Bereich zwischen der dritten und vierten horizontalen Linie sowie der ersten und vierten vertikalen Linie ausfüllt.

```js
assert.match(
  code,
  /.item5\s*?{[\s\S]*grid-area\s*?:\s*?3\s*?\/\s*?1\s*?\/\s*?4\s*?\/\s*?4\s*?;[\s\S]*}/gi
);
```

# --seed--

## --seed-contents--

```html
<style>
  .item1 {
    background: LightSkyBlue;
  }
  .item2 {
    background: LightSalmon;
  }
  .item3 {
    background: PaleTurquoise;
  }
  .item4 {
    background: LightPink;
  }

  .item5 {
    background: PaleGreen;
    /* Only change code below this line */


    /* Only change code above this line */
  }

  .container {
    font-size: 40px;
    min-height: 300px;
    width: 100%;
    background: LightGray;
    display: grid;
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: 1fr 1fr 1fr;
    grid-gap: 10px;
  }
</style>

<div class="container">
  <div class="item1">1</div>
  <div class="item2">2</div>
  <div class="item3">3</div>
  <div class="item4">4</div>
  <div class="item5">5</div>
</div>
```

# --solutions--

```html
<style>
  .item5 {
    grid-area: 3/1/4/4;
  }
</style>
```
