---
id: a77dbc43c33f39daa4429b4f
title: Побудова функції перевірки булевого значення
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

У цій лабораторній роботі ви створите функцію, яка перевіряє, чи є значення булевим примітивом.

Булеві примітиви — це `true` та `false`.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з назвою `booWho`, яка приймає один аргумент.
1. Якщо отриманий аргумент є булевим примітивом, функція має повертати `true`.
1. Якщо аргумент — будь-яке інше значення, функція має повертати `false`.

# --hints--

У вас має бути `booWho` функція.

```js
assert.isFunction(booWho);
```

`booWho(true)` має повертати `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` має повертати `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1, 2, 3])` має повертати `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` має повертати `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({ "a": 1 })` має повертати `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` має повертати `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` має повертати `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` має повертати `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` має повертати `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` має повертати `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
