---
id: 5f332b23c2045fb843337579
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

カフェが `2020` に創業されたことを訪問者に知らせるために、`p` 要素の下にテキスト `h1` を含む `Est. 2020` 要素を追加してください。

# --hints--

開く `<p>` タグがあるはずです。

```js
assert.match(code, /<p>/i);
```

閉じる `</p>` タグがあるはずです。

```js
assert.match(code, /<\/p>/i);
```

既存の `h1` 要素は変更しないでください。終了タグを削除していないことを確認してください。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

`p` 要素は `h1` 要素の下にあるはずです。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

`p` 要素にはテキスト `Est. 2020` が含まれているはずです。

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
