---
id: 5f332b23c2045fb843337579
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

To let visitors know the cafe was founded in `2020`, add a `p` element below the `h1` element with the text `Est. 2020` にしてください。

# --hints--

`<p>` の開始タグが 1 つ必要です。

```js
assert.match(code, /<p>/i);
```

終了タグ `</p>` が 1 つ必要です。

```js
assert.match(code, /<\/p>/i);
```

既存の `h1` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

`p` 要素は `h1` 要素の下に置く必要があります。

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

`p` 要素は `Est. 2020` というテキストをもつ必要があります。

```js
assert.equal(document.querySelector("p")?.innerText.trim(), "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--

--fcc-editable-region--
    </main>
  </body>
</html>
```

