---
id: 5f33294a6af5e9188dbdb8f3
title: ステップ 1
challengeType: 0
dashedName: step-1
---

# --description--

このワークショップでは、カフェのメニューを作りながらCSS（カスケーディングスタイルシート）の基本を練習します。

まずはメニューの内容を追加しましょう。既存の`main`要素の中に`body`要素を追加してください。ここには最終的にカフェで提供されるコーヒーやデザートの価格情報が入ります。

# --hints--

コードには`<main>`の開始タグが必要です。

```js
assert.match(code, /<main>/i);
```

コードには`</main>`の終了タグが必要です。

```js
assert.match(code, /<\/main>/i);
```

`body`要素は変更しないでください。終了タグを誤って削除しないように注意しましょう。

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

`main`タグは`body`タグの中に入っている必要があります。

```js
const main = document.querySelector('main');
assert.equal(main.parentElement.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
