---
id: 5f344fad8bf01691e71a30eb
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

ここまでは、コンテンツの表示や外観をコントロールする方法は限られていました。 見た目を変更できるようにするため、`head` 要素の中に `style` 要素を追加してください。

# --hints--

`<style>` の開始タグが 1 つ必要です。

```js
assert.match(code, /<style\s*>/i);
```

終了タグ `</style>` が 1 つ必要です。

```js
assert.match(code, /<\/style\s*>/);
```

`style` 要素は `head` 要素の中にネストされている必要があります。

```js
assert.match(code, /<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
