---
id: 5f344f9c805cd193c33d829c
title: ステップ 7
challengeType: 0
dashedName: step-7
---

# --description--

前のレッスンでは、次のように`CSS`のプロパティと値を追加する方法を学びました。

```css
element {
 property: value;
}
```

`h1`要素の内容を中央に配置するには、その`text-align`プロパティを`center`の値に設定してください。

# --hints--

`h1`要素に`style`セレクターがあるはずです。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasSelector);
```

`text-align`プロパティの値は`center`である必要があります。

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert.isTrue(hasTextAlign);
```

`h1`セレクターは`text-align`プロパティを`center`に設定する必要があります。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
