---
id: 5f3477cb303c5cb61b43aa9b
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

テキストは再び中央揃えになっているので、CSSファイルへのリンクは正しく機能しています。`background-color`要素の`brown`プロパティを`body`に変更する別のスタイルをファイルに追加してください。

# --hints--

`body`セレクターを使うべきです。

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert.exists(hasBody);
```

`background-color`プロパティを`brown`に設定してください。

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert.isTrue(hasBackground);
```

`body`要素は`brown`の背景を持つべきです。

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert.equal(bodyBackground, 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
h1, h2, p {
  text-align: center;
}

```

