---
id: 5f3477cb2e27333b1ab2b955
title: Step 12
challengeType: 0
dashedName: step-12
---

# --description--

スタイルを再度適用するためには `styles.css` ファイルをリンクする必要があります。 `head` 要素内に、`link` 要素追加してください。 `rel` 属性の値は `"stylesheet"`、`href` 属性の値は `"styles.css"` に設定してください。

**Note:** The `link` element is a void element, which means it doesn't have a closing tag. Void elements should be written as `<link>` rather than `<link></link>`.

# --hints--

コードには `link` 要素が必要です。

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

`link` 要素は空要素なので、終了タグ `</link>` は必要ありません。

```js
assert.notMatch(code, /<\/link>/i);
```

既存の `head` 要素を変更しないでください。 終了タグを削除していないか確認してください。

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

`link` 要素が 1 つだけ必要です。

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

`link` 要素は `head` 要素内に置く必要があります。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link` 要素には `rel` 属性があり、その値が `stylesheet` に設定されている必要があります。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link` 要素には `href` 属性があり、値が `styles.css` に設定されている必要があります。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--

--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
