---
id: 5f3477ae34c1239cafe128be
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

同じスタイルを持つ型セレクターが3つあります。セレクターのリストを作成することで、多くの要素に同じスタイルのグループを追加できます。各セレクターはカンマで区切ります。

```css
selector1, selector2 {
  property: value;
}
```

既存の3つの型セレクターを削除し、`h1`、`h2`、`p` 要素のテキストを中央揃えにする1つのセレクターリストに置き換えてください。

# --hints--

3つの要素すべてに対して、単一の型セレクターを使う必要があります。`h1, h2, p` の順番を必ず守ってください。

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert.exists(hasSelector);
```

`style` 要素にはセレクターを1つだけにしてください。

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert.equal(selectors.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
--fcc-editable-region--
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

