---
id: 5f3477ae9675db8bb7655b30
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

前のステップでは、<dfn>型セレクター</dfn>を使って`h1`要素にスタイルを適用しました。既存の`h2`要素に、それぞれ`p`と`style`要素の内容を中央揃えにする新しい型セレクターを追加してください。

# --hints--

既存の`h1`セレクターは変更しないでください。

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert.exists(hasH1);
```

新しい`style`タグを追加しないでください。新しいCSSルールは既存の`style`タグに追加してください。

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

新しい`h2`セレクターを追加してください。

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert.exists(hasH2);
```

新しい`p`セレクターを追加してください。

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert.exists(hasP);
```

`h1`要素は`text-align`が`center`である必要があります。

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert.equal(h1TextAlign, 'center');
```

`h2`要素は`text-align`が`center`である必要があります。

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert.equal(h2TextAlign, 'center');
```

`p`要素は`text-align`が`center`である必要があります。

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert.equal(pTextAlign, 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1 {
        text-align: center;
      }
--fcc-editable-region--
      
--fcc-editable-region--
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

