---
id: 5f3477cb2e27333b1ab2b955
title: ステップ 12
challengeType: 0
dashedName: step-12
---

# --description--

今度は`styles.css`ファイルをリンクして、スタイルを再び適用できるようにします。`head`要素の中に`link`要素を追加してください。`rel`属性に`"stylesheet"`の値を、`href`属性に`"styles.css"`の値を付けてください。

**注意:** `link`要素は空要素で、終了タグがありません。空要素は`<link>`のように書き、`<link></link>`のようには書かないでください。

# --hints--

コードには`link`要素が必要です。

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

`link`要素は空要素なので、終了タグ`</link>`は不要です。

```js
assert.notMatch(code, /<\/link>/i);
```

既存の`head`要素は変更しないでください。終了タグを削除していないことを確認してください。

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

`link`要素は1つだけ必要です。

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

`link`要素は`head`要素の中にある必要があります。

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

`link`要素には`rel`属性があり、その値は`stylesheet`である必要があります。

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

`link`要素には`href`属性があり、その値は`styles.css`である必要があります。

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
