---
id: 5f3477aefa51bfc29327200b
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

`style` タグ内に CSS を書いて、3つの要素にスタイルを適用しました。これは動作しますが、スタイルがもっと増えるので、すべてのスタイルを別のファイルにまとめてリンクするのが最適です。

別の `styles.css` ファイルが用意されています。エディター上部のタブでファイルを切り替えられます。

まず、作成したスタイルを `styles.css` ファイルに書き直してください。開く `style` タグと閉じる `style` タグは含めないように注意してください。

# --hints--

`styles.css` ファイルには `h1, h2, p` 型セレクターを使ってください。

```js
assert.match(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, ""), __helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/));
```
 
セレクターは `text-align` プロパティを `center` に設定する必要があります。

```js
assert.match(code.slice(code.indexOf('</html>')), /text-align:\s*center;?/);
```

コードに `<style>` タグを含めないでください。

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
