---
id: 5f344fad8bf01691e71a30eb
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

これまでは、コンテンツの表示や見た目をあまり自由にコントロールできませんでした。これを変えるために、`style` 要素の中に `head` 要素を追加してください。

# --hints--

コードには `style` 要素を含めてください。

```js
assert.exists(document.querySelector('style:not(.fcc-hide-header)'));
```

コードには閉じる `</style>` タグが必要です。

```js
assert.match(code, /<\/style\s*>/i);
```

`style` 要素は `head` 要素の中にネストされている必要があります。

```js
assert.exists(document.querySelector('head > style:not(.fcc-hide-header)'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
