---
id: 5f344fc1520b6719f2e35605
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

メニューにはコーヒー用とデザート用の2つのセクションがあります。すべてのコーヒーを入れる場所を作るために、`section` 要素の中に `main` 要素を追加してください。

# --hints--

`<section>` の開始タグがあるはずです。

```js
assert.match(code, /<section\s*>/i);
```

`</section>` の終了タグがあるはずです。

```js
assert.match(code, /<\/section\s*>/i);
```

既存の `main` 要素は変更しないでください。終了タグを削除していないか確認してください。

```js
assert.lengthOf(document.querySelectorAll('main'),1);
```

`section` 要素は `main` 要素の中にある必要があります。

```js
const main = document.querySelector('main');
const section = document.querySelector('section');
assert.strictEqual(section.parentElement, main);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
