---
id: 5f356ed656a336993abd9f7c
title: Step 21
challengeType: 0
dashedName: step-21
---

# --description--

次に、`#menu`を水平方向に中央揃えにしたいです。これは`margin-left`と`margin-right`のプロパティを`auto`に設定することでできます。マージンは要素の周りの見えないスペースのようなものだと考えてください。この2つのマージンプロパティを使って、`#menu`要素の中で`body`要素を中央に配置してください。

# --hints--

`margin-left`プロパティを`auto`に設定するべきです。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert.isTrue(hasMargin);
```

`margin-right`プロパティを`auto`に設定するべきです。

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert.isTrue(hasMargin);
```

`margin-left`の`margin-right`と`#menu`プロパティを`auto`に設定してください。

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert.equal(divMarginRight, 'auto');
assert.equal(divMarginLeft, 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
  width: 80%;
  background-color: burlywood;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
