---
id: 5f356ed60785e1f3e9850b6e
title: ステップ 20
challengeType: 0
dashedName: step-20
---

# --description--

`#menu` 要素内のテキストが中央揃えになっているのが簡単にわかります。現在、`#menu` 要素の幅はピクセル（`px`）で指定されています。

`width` プロパティの値を `80%` に変更して、親要素（`80%`）の幅の `body` にしてください。

# --hints--

`width` プロパティを `80%` に設定してください。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert.isTrue(hasWidth);
```

`width` プロパティを `300px` にしてはいけません。

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert.isFalse(hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

#menu {
--fcc-editable-region--
  width: 300px;
--fcc-editable-region--
  background-color: burlywood;
}
```
