---
id: 5f3c866dbf362f99b9a0c6d0
title: Step 33
challengeType: 0
dashedName: step-33
---

# --description--

`p` 要素は、クラス属性が `article` の `item` 要素にネストされています。クラス名が `p` の要素内のどこにネストされている `item` 要素すべてに対して、次のようにスタイルを適用できます。

```css
.item p { }
```

上記のセレクターを使って、`display` 要素がよりインライン要素のように振る舞うように、値が `inline-block` の `p` プロパティを追加してください。

# --hints--

`.item p` セレクターを使うべきです。

```js
const itemPStyle = new __helpers.CSSHelp(document).getStyle('.item p');
assert.exists(itemPStyle);
```

`display` プロパティを `inline-block` に設定してください。

```js
const hasDisplay = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.display === 'inline-block');
assert.isTrue(hasDisplay);
```

`.item p` セレクターは `display` プロパティを `inline-block` に設定するべきです。

```js
const itemPDisplay = new __helpers.CSSHelp(document).getStyle('.item p')?.getPropertyValue('display');
assert.equal(itemPDisplay, 'inline-block');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

--fcc-editable-region--

--fcc-editable-region--

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
