---
id: 5f3c866d5414453fc2d7b480
title: ステップ 27
challengeType: 0
dashedName: step-27
---

# --description--

既存の coffee/price ペアの下に、各 `article` 要素の中に二つのネストされた `p` 要素を使って、次のコーヒーと価格を追加してください。

```md
Caramel Macchiato 3.75
Pumpkin Spice 3.50
Hazelnut 4.00
Mocha 4.50
```

これまでと同様に、最初の `p` 要素のテキストにはコーヒーのフレーバーを、二番目の `p` 要素のテキストには価格を入れてください。

# --hints--

`article` 要素が5つあるはずです。

```js
assert.lengthOf(document.querySelectorAll('article'),5);
```

各 `article` 要素は二つの `p` 要素を持つべきです。

```js
const articles = document.querySelectorAll('article');
assert.lengthOf(articles?.[0]?.children, 2);
assert.lengthOf(articles?.[1]?.children, 2);
assert.lengthOf(articles?.[2]?.children, 2);
assert.lengthOf(articles?.[3]?.children, 2);
assert.lengthOf(articles?.[4]?.children, 2);
```

最初の `article` 要素は、テキストが `p` と `French Vanilla` の `3.00` 要素を持つべきです。

```js
const children = document.querySelector('article')?.children;
assert.match(children?.[0]?.innerText.trim(), /French Vanilla/i);
assert.match(children?.[1]?.innerText.trim(), /3\.00/i);
```

二番目の `article` 要素は、テキストが `p` と `Caramel Macchiato` の `3.75` 要素を持つべきです。

```js
const children = document.querySelectorAll('article')?.[1]?.children;
assert.match(children?.[0]?.innerText.trim(), /Caramel Macchiato/i);
assert.match(children?.[1]?.innerText.trim(), /3\.75/i);
```

三番目の `article` 要素は、テキストが `p` と `Pumpkin Spice` の `3.50` 要素を持つべきです。

```js
const children = document.querySelectorAll('article')?.[2]?.children;
assert.match(children?.[0]?.innerText.trim(), /Pumpkin Spice/i);
assert.match(children?.[1]?.innerText.trim(), /3\.50/i);
```

四番目の `article` 要素は、テキストが `p` と `Hazelnut` の `4.00` 要素を持つべきです。

```js
const children = document.querySelectorAll('article')?.[3]?.children;
assert.match(children?.[0]?.innerText.trim(), /Hazelnut/i);
assert.match(children?.[1]?.innerText.trim(), /4\.00/i);
```

五番目の `article` 要素は、テキストが `p` と `Mocha` の `4.50` 要素を持つべきです。

```js
const children = document.querySelectorAll('article')?.[4]?.children;
assert.match(children?.[0]?.innerText.trim(), /Mocha/i);
assert.match(children?.[1]?.innerText.trim(), /4\.50/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article>
            <p>French Vanilla</p>
            <p>3.00</p>
          </article>
--fcc-editable-region--
          
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
