---
id: 66f2836c459cfb16ae76f24f
title: Створіть генератор пірамід
challengeType: 26
dashedName: lab-pyramid-generator
---

# --description--

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати функцію під назвою `pyramid`, яка приймає три аргументи.
1. Першим аргументом має бути рядок, який представляє символ, що повторюється у піраміді.
1. Другим аргументом має бути ціле число, яке представляє кількість рядків у піраміді.
1. Третім аргументом має бути булеве значення.
1. Функція `pyramid` має повернути рядок, де символ повторюється та впорядковується так, щоб утворити піраміду, вершина якої спрямована вгору, якщо третім аргументом є `false`.
1. Якщо третім аргументом є `true`, то вершина піраміди має бути спрямована донизу.
1. Рядок вершини повинен мати один символ, а кожен інший рядок повинен мати на два символи більше, ніж попередній.
1. Each row should start with a number of spaces sufficient to put the center character of each row in the same column and there should not be any spaces at the end of each row.
1. Піраміда має починатися та закінчуватися символом нового рядка.

Наприклад, при виклику `pyramid("o", 4, false)` має бути такий результат:

```js

   o
  ooo
 ooooo
ooooooo

```

# --hints--

Ви повинні мати функцію під назвою `pyramid`.

```js
assert.isFunction(pyramid);
```

Функція `pyramid` повинна мати три параметри.

```js
assert.lengthOf(pyramid, 3);
```

`pyramid("o", 4, false)` має повернути `"\n   o\n  ooo\n ooooo\nooooooo\n"`.

```js
assert.equal(pyramid("o", 4, false), "\n   o\n  ooo\n ooooo\nooooooo\n")
```

`pyramid("p", 5, true)` має повернути `"\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n"`.

```js
assert.equal(pyramid("p", 5, true), "\nppppppppp\n ppppppp\n  ppppp\n   ppp\n    p\n")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function pyramid(char, count, isInverted) {
    const rows = []
    for (let i = 1; i <= count; i++) {
        if (isInverted) {
            rows.unshift(" ".repeat(count - i) + char.repeat(2 * i - 1))
        } else {
            rows.push(" ".repeat(count - i) + char.repeat(2 * i - 1))
        }        
    }
    return "\n" + rows.join("\n") + "\n";
}

console.log(pyramid("#", 10, false))

```
