---
id: 6140c7e645d8e905819f1dd4
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

適切な `charset` を指定する `meta` 要素、`title` 要素、および `./styles.css` ファイルを指定する `link` 要素を追加してください。

`title` を `Ferris Wheel` に設定してください。

# --hints--

コードに `DOCTYPE` 宣言が含まれている必要があります。

```js
assert.match(code, /<!DOCTYPE/gi);
```

`DOCTYPE` 宣言の後に半角スペースを入れる必要があります。

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

ドキュメントタイプが `html` であると定義する必要があります。

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

タイプの後ろに `>` を入力して、`DOCTYPE` 宣言を閉じる必要があります。

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

`html` 要素は `lang` 属性が `en` の開始タグをもつ必要があります。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

`html` 要素には終了タグが必要です。

```js
assert.match(code, /<\/html\s*>/);
```

`DOCTYPE` 宣言は HTML の先頭にある必要があります。

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

`head` の開始タグが 1 つ必要です。

```js
assert.match(code, /<head\s*>/i);
```

`head` の終了タグが 1 つ必要です。

```js
assert.match(code, /<\/head\s*>/i);
```

`body` の開始タグが 1 つ必要です。

```js
assert.match(code, /<body\s*>/i);
```

`body` の終了タグが 1 つ必要です。

```js
assert.match(code, /<\/body\s*>/i);
```

`body` 要素を `head` 要素の直後に置く必要があります。

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

`head` 要素と `body` 要素は兄弟要素である必要があります。

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

`head` 要素は `html` 要素の内側の、最初の要素である必要があります。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

`head` 要素は `html` 要素の中にある必要があります。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 要素は `html` 要素の中にある必要があります。

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

`meta` 要素は `head` 要素内に置く必要があります。

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

`link` 要素が 1 つだけ必要です。

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
