---
id: 68a47dcdbce12193caf77012
title: implementare un cercatore di indice ordinato
challengeType: 26
dashedName: lab-sorted-index-finder
---

# --description--

In questo laboratorio creerai una funzione che restituisce l'indice più basso in cui un valore dovrebbe essere inserito in un array una volta ordinato in ordine crescente.

**Obiettivo**: soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `getIndexToIns` che prende due argomenti: un array e un numero.  
2. Dovresti usare il metodo `sort` per ordinare l'array in ordine crescente.
3. La tua funzione `getIndexToIns` dovrebbe restituire l'indice più basso in cui il numero dovrebbe essere inserito usando il metodo `findIndex`.
4. La tua funzione `getIndexToIns` dovrebbe sempre restituire un numero.

**Suggerimento:**  
Il metodo `findIndex` è un metodo integrato degli array in JavaScript. Prende una funzione di callback e restituisce l'indice del primo elemento che soddisfa la condizione. Sia `findIndex` che `sort` sono funzioni di ordine superiore.

**Esempi:**

- `getIndexToIns([1, 2, 3, 4], 1.5)` dovrebbe restituire `1` perché `1.5` è maggiore di `1` (indice `0`) e minore di `2` (indice `1`).
- `getIndexToIns([20, 3, 5], 19)` dovrebbe restituire `2` perché dopo l'ordinamento a `[3, 5, 20]`, `19` è minore di `20` (indice `2`) e maggiore di `5` (indice `1`).


# --hints--

Dovresti avere una funzione `getIndexToIns`.

```js
assert.isFunction(getIndexToIns);
```

`getIndexToIns` dovrebbe sempre restituire un numero.

```js
assert.isNumber(getIndexToIns([10, 20, 30, 40, 50], 35));
```

La tua funzione dovrebbe usare il metodo `sort`.

```js
assert.match(getIndexToIns.toString(), /\.sort\(/);
```


La tua funzione dovrebbe usare il metodo `findIndex`.

```js
assert.match(getIndexToIns.toString(), /\.findIndex\(/);
```

`getIndexToIns([10, 20, 30, 40, 50], 35)` dovrebbe restituire `3`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 35), 3);
```

`getIndexToIns([10, 20, 30, 40, 50], 30)` dovrebbe restituire `2`.

```js
assert.strictEqual(getIndexToIns([10, 20, 30, 40, 50], 30), 2);
```

`getIndexToIns([40, 60], 50)` dovrebbe restituire `1`.

```js
assert.strictEqual(getIndexToIns([40, 60], 50), 1);
```

`getIndexToIns([3, 10, 5], 3)` dovrebbe restituire `0`.

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 3), 0);
```

`getIndexToIns([5, 3, 20, 3], 5)` dovrebbe restituire `2`.

```js
assert.strictEqual(getIndexToIns([5, 3, 20, 3], 5), 2);
```

`getIndexToIns([2, 20, 10], 19)` dovrebbe restituire `2`.

```js
assert.strictEqual(getIndexToIns([2, 20, 10], 19), 2);
```

`getIndexToIns([3, 10, 5], 11)` dovrebbe restituire `3`

```js
assert.strictEqual(getIndexToIns([3, 10, 5], 11), 3);
```

`getIndexToIns([], 5)` dovrebbe restituire `0`

```js
assert.strictEqual(getIndexToIns([], 5), 0);
```

# --seed--

## --seed-contents--

```js
``` 

# --solutions--

```js
function getIndexToIns(arr, num) {
  const sorted = [...arr].sort((a, b) => a - b);
  const index = sorted.findIndex(el => el >= num);

  return index === -1 ? sorted.length : index;
}
```
