---
id: 69860157f7fce39f9df399a3
title: "implementare l'algoritmo delle n regine"
challengeType: 26
dashedName: implement-the-n-queens-algorithm
---

# --description--

Il problema delle N regine ti chiede di posizionare N regine su una scacchiera N×N in modo che nessuna regina attacchi un'altra (nessuna condivide la stessa riga, colonna o diagonale).

Per esempio, se c'è una scacchiera 4x4, una disposizione valida è:

```md
[1, 3, 0, 2]
```

Questo significa che nella riga 0 la regina è posizionata nella colonna 1; nella riga 1 la regina è posizionata nella colonna 3; nella riga 2 la regina è posizionata nella colonna 0; e nella riga 3 la regina è posizionata nella colonna 2.

Visivamente, questa disposizione appare così:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

Dove `Q` rappresenta una regina e `.` rappresenta una casella vuota.

In questo laboratorio, implementerai il risolutore del problema delle N regine usando l'approccio della ricerca in profondità.


**Obiettivo**: Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `dfsNQueens`.
2. La funzione dovrebbe accettare esattamente un argomento: un intero `n`.
3. Se `n` è minore di `1`, la funzione dovrebbe restituire un array vuoto (`[]`).
4. La funzione dovrebbe restituire una lista di soluzioni; ogni soluzione è a sua volta una lista di lunghezza `n`, dove l'elemento all'indice `i` è l'indice di colonna (basato su zero) della regina nella riga `i`.

# --hints--

Dovresti avere una funzione chiamata `dfsNQueens` che prende un argomento.

```js
assert.isFunction(dfsNQueens);
assert.lengthOf(dfsNQueens, 1);
```

Se `n` è minore di `1`, la funzione dovrebbe restituire un array vuoto.

```js
assert.sameDeepOrderedMembers(dfsNQueens(0), []);
assert.sameDeepOrderedMembers(dfsNQueens(-1), []);
assert.sameDeepOrderedMembers(dfsNQueens(-5), []);
```

La funzione dovrebbe restituire un array di soluzioni, dove ogni soluzione è un array di lunghezza `n`.

```js
const output = dfsNQueens(4);
assert.isArray(output);
output.forEach(solution => {
    assert.isArray(solution);
    assert.lengthOf(solution, 4);
})
```

`dfsNQueens(1)` dovrebbe restituire `[[0]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(1), [[0]]);
```

`dfsNQueens(2)` dovrebbe restituire `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(2), []);
```

`dfsNQueens(3)` dovrebbe restituire `[]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(3), []);
```

`dfsNQueens(4)` dovrebbe restituire `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(4), [[1, 3, 0, 2], [2, 0, 3, 1]]);
```

`dfsNQueens(5)` dovrebbe restituire `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
assert.sameDeepOrderedMembers(dfsNQueens(5), [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]);
```

`dfsNQueens(5).length` dovrebbe essere `10`.

```js
assert.lengthOf(dfsNQueens(5), 10);
```

`dfsNQueens(8).length` dovrebbe essere `92`.

```js
assert.lengthOf(dfsNQueens(8), 92);
```

`dfsNQueens` dovrebbe restituire il risultato corretto.

```js
assert.lengthOf(dfsNQueens(7), 40);
assert.sameDeepOrderedMembers(dfsNQueens(6), [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfsNQueens(n) {
  if (n < 1) return [];

  const results = [];

  function isSafe(queens, row, col) {
    for (let r = 0; r < queens.length; r++) {
      const c = queens[r];
      if (c === col || Math.abs(row - r) === Math.abs(col - c)) {
        return false;
      }
    }
    return true;
  }

  function dfs(queens) {
    const row = queens.length;
    if (row === n) {
      results.push(queens.slice());
      return;
    }

    for (let col = 0; col < n; col++) {
      if (isSafe(queens, row, col)) {
        queens.push(col);
        dfs(queens);
        queens.pop();
      }
    }
  }

  dfs([]);
  return results;
}
```
