---
id: 587d7788367417b2b2512aa2
title: rendi più semplice la navigazione con lettore di schermo usando il landmark nav
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVwWSv'
forumTopicId: 301024
dashedName: make-screen-reader-navigation-easier-with-the-nav-landmark
---

# --description--

L'elemento `nav` è un altro elemento HTML5 con la funzionalità landmark incorporata per facilitare la navigazione con il lettore di schermo. Questo tag serve a racchiudere i collegamenti principali di navigazione nella tua pagina.

Se ci sono collegamenti ripetuti del sito in fondo alla pagina, non è necessario contrassegnarli anche con un tag `nav`. Usare un `footer` (trattato nella sfida successiva) è sufficiente.

# --instructions--

Camper Cat ha inserito i collegamenti di navigazione in cima alla sua pagina di allenamento, ma li ha racchiusi in un `div`. Cambia il `div` in un tag `nav` per migliorare l'accessibilità della sua pagina.

# --hints--

Il tuo codice dovrebbe avere un solo tag `nav`.

```js
assert.lengthOf(document.querySelectorAll('nav') , 1);
```

I tuoi tag `nav` dovrebbero racchiudere `ul` e i suoi elementi di lista.

```js
const nav = document.querySelector('nav');
const children = nav?.querySelectorAll(`:scope ${'ul'}`);
assert.lengthOf(children,1);
```

Il tuo codice non dovrebbe contenere tag `div`.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Il tuo elemento `nav` dovrebbe avere un tag di chiusura.

```js
assert.isTrue(code.match(/<\/nav>/g)?.length === code.match(/<nav>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <div>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </div>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
