---
id: 587d778a367417b2b2512aa5
title: "migliorare l'accessibilità dei grafici con l'elemento figure"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5 ha introdotto l'elemento `figure` e il relativo `figcaption`. Usati insieme, questi elementi racchiudono una rappresentazione visiva (come un'immagine, un diagramma o un grafico) insieme al suo sottotitolo. Raggruppare questi elementi insieme offre un doppio vantaggio in termini di accessibilità, raggruppando semanticamente contenuti correlati e fornendo un'alternativa testuale che spiega il `figure`.

Per le visualizzazioni dei dati come i grafici, il sottotitolo può essere usato per annotare brevemente le tendenze o le conclusioni per gli utenti con disabilità visiva. Un'altra sfida spiega come spostare una versione tabellare dei dati del grafico fuori schermo (usando CSS) per gli utenti di lettori di schermo.

Ecco un esempio: nota che il `figcaption` va all'interno dei tag `figure` e può essere combinato con altri elementi:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat sta lavorando duramente per creare un grafico a barre impilate che mostra il tempo settimanale da dedicare all'allenamento in furtività, combattimento e armi. Aiutalo a strutturare meglio la sua pagina cambiando il tag `div` che ha usato in un tag `figure`, e il tag `p` che circonda il sottotitolo in un tag `figcaption`.

# --hints--

Il tuo codice dovrebbe avere un tag `figure`.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

Il tuo codice dovrebbe avere un tag `figcaption`.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

Il tuo codice non dovrebbe avere tag `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Il tuo codice non dovrebbe avere tag `p`.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

Il `figcaption` dovrebbe essere un figlio del tag `figure`.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

Il tuo elemento `figure` dovrebbe avere un tag di chiusura.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
