---
id: 587d774e367417b2b2512aa0
title: "Inserisci il contenuto nell'elemento article"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` è un altro dei nuovi elementi HTML5 che aggiungono significato semantico al tuo markup. `article` è un elemento di sezione e si usa per racchiudere contenuti indipendenti e autonomi. Il tag funziona bene con articoli di blog, post di forum o articoli di notizie.

Stabilire se un contenuto può stare da solo è spesso una questione di giudizio, ma puoi usare un paio di test semplici. Chiediti se, rimuovendo tutto il contesto circostante, quel contenuto avrebbe ancora senso. Allo stesso modo, per il testo, il contenuto reggerebbe se fosse in un feed RSS?

Ricorda che chi usa tecnologie assistive si affida a un markup organizzato e semanticamente significativo per comprendere meglio il tuo lavoro.

**Nota:** L'elemento `section` è anch’esso nuovo con HTML5 e ha un significato semantico leggermente diverso rispetto a `article`. Un `article` è per contenuti autonomi, mentre un `section` serve a raggruppare contenuti tematicamente correlati. Possono essere usati l’uno dentro l’altro, se necessario. Per esempio, se un libro è il `article`, allora ogni capitolo è un `section`. Quando non c’è relazione tra gruppi di contenuti, allora usa un `div`.

`<div>` - raggruppa contenuti  
`<section>` - raggruppa contenuti correlati  
`<article>` - raggruppa contenuti indipendenti e autonomi  

# --instructions--

Camper Cat ha usato tag `article` per racchiudere i post nella sua pagina blog, ma ha dimenticato di usarli attorno al primo. Cambia il tag `div` in modo che usi un tag `article` invece.

# --hints--

Il tuo codice dovrebbe avere tre tag `article`.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

Il tuo codice non dovrebbe contenere tag `div`.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
