---
id: 587d778f367417b2b2512aac
title: Evita problemi di daltonismo usando un contrasto sufficiente
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

Il colore è una parte importante del design visivo, ma il suo uso introduce due problemi di accessibilità. Primo, il colore da solo non dovrebbe essere usato come unico modo per trasmettere informazioni importanti perché chi usa un lettore di schermo non lo vede. Secondo, i colori di primo piano e sfondo devono avere un contrasto sufficiente affinché gli utenti daltonici possano distinguerli.

Le sfide precedenti hanno trattato l’uso di alternative testuali per risolvere il primo problema. L’ultima sfida ha introdotto strumenti per controllare il contrasto e aiutare con il secondo. Il rapporto di contrasto raccomandato dalle WCAG di 4,5:1 si applica sia all’uso del colore sia alle combinazioni in scala di grigi.

Gli utenti daltonici hanno difficoltà a distinguere alcuni colori da altri — di solito per la tonalità ma a volte anche per la luminosità. Ricorderai che il rapporto di contrasto si calcola usando i valori di luminanza relativa (o luminosità) dei colori di primo piano e sfondo.

In pratica, il rapporto di contrasto 4,5:1 si può raggiungere scurendo (aggiungendo nero) il colore più scuro e schiarendo (aggiungendo bianco) il colore più chiaro. Le tonalità più scure nella ruota dei colori sono considerate sfumature di blu, viola, magenta e rosso, mentre i colori più chiari sono arancioni, gialli, verdi e verde-blu.

# --instructions--

Camper Cat sta sperimentando l’uso del colore per il testo e lo sfondo del suo blog, ma la sua combinazione attuale di un verde `background-color` con testo color marrone `color` ha un rapporto di contrasto di 2,5:1. Puoi facilmente regolare la luminosità dei colori dato che li ha dichiarati usando la proprietà CSS `hsl()` (che indica tonalità, saturazione, luminosità) cambiando il terzo argomento. Aumenta il valore di luminosità `background-color` dal 35% al 55% e diminuisci il valore di luminosità `color` dal 20% al 15%. Questo migliora il contrasto a 5,9:1.

# --hints--

Il tuo codice dovrebbe cambiare solo il valore di luminosità per la proprietà `color` del testo a 15%.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

Il tuo codice dovrebbe cambiare solo il valore di luminosità per la proprietà `background-color` a 55%.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
