---
id: 587d778f367417b2b2512aae
title: Dare significato ai collegamenti usando testo descrittivo
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

Gli utenti di lettori di schermo hanno varie opzioni sul tipo di contenuto che il loro dispositivo legge. Queste opzioni includono saltare agli elementi di riferimento (o oltre), passare al contenuto principale o ottenere un riepilogo della pagina dalle intestazioni. Un'altra opzione è ascoltare solo i collegamenti disponibili su una pagina.

I lettori di schermo fanno questo leggendo il testo del collegamento, cioè ciò che si trova tra i tag àncora (`a`). Avere una lista di collegamenti con testo come "clicca qui" o "leggi di più" non è utile. Invece, usa un testo breve ma descrittivo all'interno dei tag `a` per fornire più significato a questi utenti.

# --instructions--

Il testo del collegamento che Camper Cat sta usando non è molto descrittivo senza il contesto circostante. Sposta i tag àncora (`a`) in modo che racchiudano il testo `information about batteries` invece di `Click here`.

# --hints--

Il tuo codice dovrebbe spostare i tag àncora `a` dalle parole `Click here` per racchiudere le parole `information about batteries`.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

L'elemento `a` dovrebbe avere un attributo `href` con valore stringa vuota `""`.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

L'elemento `a` dovrebbe avere un tag di chiusura.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
