---
id: 587d7790367417b2b2512ab0
title: usa tabindex per aggiungere il focus da tastiera a un elemento
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

L'attributo HTML `tabindex` ha tre funzioni distinte relative al focus da tastiera di un elemento. Quando è presente su un tag, indica che l'elemento può ricevere il focus. Il valore (un numero intero positivo, negativo o zero) determina il comportamento.

Alcuni elementi, come i collegamenti e i controlli del modulo, ricevono automaticamente il focus da tastiera quando un utente naviga con il tasto tab in una pagina. L'ordine è lo stesso in cui gli elementi appaiono nel markup HTML sorgente. Questa stessa funzionalità può essere assegnata ad altri elementi, come `div`, `span` e `p`, aggiungendo loro un attributo `tabindex="0"`. Ecco un esempio:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**Nota:** Un valore `tabindex` negativo (tipicamente -1) indica che un elemento è focalizzabile, ma non raggiungibile tramite tastiera. Questo metodo viene generalmente usato per portare il focus al contenuto in modo programmato (ad esempio quando si attiva un `div` usato per una finestra pop-up), ed è fuori dall'ambito di queste sfide.

# --instructions--

Camper Cat ha creato un nuovo sondaggio per raccogliere informazioni sui suoi utenti. Sa che i campi di input ricevono automaticamente il focus da tastiera, ma vuole assicurarsi che gli utenti da tastiera si fermino sulle istruzioni mentre navigano con il tab tra gli elementi. Aggiungi un attributo `tabindex` al tag `p` e imposta il suo valore su `0`. Bonus: usare `tabindex` abilita anche la pseudo-classe CSS `:focus` sul tag `p`.

# --hints--

Il tuo codice dovrebbe aggiungere un attributo `tabindex` al tag `p` che contiene le istruzioni del modulo.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

Il tuo codice dovrebbe impostare l'attributo `tabindex` sul tag `p` con un valore di 0.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
