---
id: 587d778b367417b2b2512aa8
title: aggiungi un selettore di date accessibile
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

I moduli spesso includono il campo `input`, che può essere usato per creare diversi controlli del modulo. L'attributo `type` su questo elemento indica che tipo di elemento `input` verrà creato.

Potresti aver notato i tipi di input `text` e `submit` in sfide precedenti, e HTML5 ha introdotto un'opzione per specificare un campo `date`. A seconda del supporto del browser, un selettore di date appare nel campo `input` quando è in focus, il che rende più facile compilare un modulo per tutti gli utenti.

Per i browser più vecchi, il tipo sarà predefinito a `text`, quindi è utile mostrare agli utenti il formato data previsto nel testo `label` o `placeholder` per sicurezza.

Ecco un esempio:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat sta organizzando un torneo di Mortal Kombat e vuole chiedere ai suoi concorrenti quale data funziona meglio. Aggiungi un tag `input` con un attributo `type` di `date`, un attributo `id` di `pickdate` e un attributo `name` di `date`.

# --hints--

Il tuo codice dovrebbe aggiungere un tag `input` per il campo selettore di date.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

Il tuo tag `input` dovrebbe avere un attributo `type` con valore `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

Il tuo tag `input` dovrebbe avere un attributo `id` con valore `pickdate`.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

Il tuo tag `input` dovrebbe avere un attributo `name` con valore `date`.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
