---
id: 587d774c367417b2b2512a9c
title: "Aggiungi un'alternativa testuale alle immagini per l'accessibilità alle persone con disabilità visiva"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Probabilmente hai già visto un attributo `alt` su un tag `img` in altre sfide. Il testo `alt` descrive il contenuto dell'immagine e ne fornisce un'alternativa testuale. Un attributo `alt` è utile nei casi in cui l'immagine non si carica o non può essere vista da un utente. Anche i motori di ricerca lo usano per capire cosa contiene un'immagine e includerla nei risultati di ricerca. Ecco un esempio:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Le persone con disabilità visiva si affidano ai lettori di schermo per convertire il contenuto web in un'interfaccia audio. Non riceveranno informazioni se queste sono presentate solo visivamente. Per le immagini, i lettori di schermo possono consultare l'attributo `alt` e leggerne il contenuto per fornire informazioni chiave.

Un buon testo `alt` offre al lettore una breve descrizione dell'immagine. Dovresti sempre includere un attributo `alt` nelle tue immagini. Secondo la specifica HTML5, questo è ormai considerato obbligatorio.

# --instructions--

Camper Cat è sia un ninja della programmazione sia un vero ninja, e sta costruendo un sito per condividere le sue conoscenze. La foto del profilo che vuole usare mostra le sue abilità e dovrebbe essere apprezzata da tutti i visitatori del sito. Aggiungi un attributo `alt` nel tag `img` che spieghi che Camper Cat sta facendo karate. (L'attributo `src` dell'immagine non punta a un file reale, quindi dovresti vedere il testo `alt` nella visualizzazione.)

# --hints--

Il tuo tag `img` dovrebbe avere un attributo `alt` e non dovrebbe essere vuoto.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
