---
id: 587d7790367417b2b2512aaf
title: rendi i collegamenti navigabili con i tasti di accesso HTML
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML offre l'attributo `accesskey` per specificare un tasto di scelta rapida per attivare o portare il focus su un elemento. Aggiungere un attributo `accesskey` può rendere la navigazione più efficiente per chi usa solo la tastiera.

HTML5 permette di usare questo attributo su qualsiasi elemento, ma è particolarmente utile quando viene usato con elementi interattivi. Questo include collegamenti, pulsanti e controlli del modulo.

Ecco un esempio:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat vuole che i collegamenti intorno ai due titoli degli articoli del blog abbiano tasti di scelta rapida in modo che gli utenti del suo sito possano navigare rapidamente alla storia completa. Aggiungi un attributo `accesskey` a entrambi i collegamenti e imposta il primo su `g` (per Garfield) e il secondo su `c` (per Chuck Norris).

# --hints--

Il tuo codice dovrebbe aggiungere un attributo `accesskey` al tag `a` con il valore `id` `first`.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

Il tuo codice dovrebbe aggiungere un attributo `accesskey` al tag `a` con il valore `id` `second`.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

Il tuo codice dovrebbe impostare l'attributo `accesskey` sul tag `a` con il valore `id` `first` su `g`. Nota che fa differenza tra maiuscole e minuscole.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

Il tuo codice dovrebbe impostare l'attributo `accesskey` sul tag `a` con il valore `id` `second` su `c`. Nota che fa differenza tra maiuscole e minuscole.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
