---
id: 587d774c367417b2b2512a9d
title: Quando lasciare vuoto il testo alternativo
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

Nell'ultima sfida, hai imparato che includere un attributo `alt` quando si usano tag `img` è obbligatorio. Tuttavia, a volte le immagini sono raggruppate con un sottotitolo che le descrive già, oppure sono usate solo per decorazione. In questi casi, il testo `alt` può sembrare ridondante o non necessario.

Quando un'immagine è già spiegata con del testo o non aggiunge significato a una pagina, il `img` necessita comunque di un attributo `alt`, ma può essere impostato su una stringa vuota. Ecco un esempio:

```html
<img src="visualDecoration.jpeg" alt="">
```

Le immagini di sfondo rientrano di solito nella categoria "decorativa". Tuttavia, sono tipicamente applicate con regole CSS e quindi non fanno parte del markup che i lettori di schermo elaborano.

**Nota:** Per le immagini con un sottotitolo, potresti comunque voler includere un testo `alt` perché aiuta i motori di ricerca a catalogare il contenuto dell'immagine.

# --instructions--

Camper Cat ha scritto il codice di una pagina scheletro per la parte blog del suo sito. Ha intenzione di aggiungere una pausa visiva tra i suoi due articoli con un'immagine decorativa di una spada samurai. Aggiungi un attributo `alt` al tag `img` e impostalo su una stringa vuota. (Nota che l'immagine `src` non punta a un file reale — non preoccuparti se non si vedono spade nella visualizzazione.)

# --hints--

Il tuo tag `img` dovrebbe avere un attributo `alt`.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

L'attributo `alt` dovrebbe essere impostato su una stringa vuota.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
