---
id: 587d7789367417b2b2512aa4
title: "migliorare l'accessibilità dei contenuti audio con l'elemento audio"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVkcZ'
forumTopicId: 301014
dashedName: improve-accessibility-of-audio-content-with-the-audio-element
---

# --description--

L'elemento `audio` di HTML5 dà un significato semantico quando racchiude contenuti sonori o flussi audio nel tuo markup. I contenuti audio hanno anche bisogno di un'alternativa testuale per essere accessibili alle persone sorde o con problemi di udito. Questo può essere fatto con un testo vicino sulla pagina o un collegamento a una trascrizione.

Il tag `audio` supporta l'attributo `controls`. Questo mostra i controlli predefiniti del browser per riproduzione, pausa e altri, e supporta la funzionalità da tastiera. Si tratta di un attributo booleano, quindi non necessita di un valore: la sua presenza sul tag attiva l'impostazione.

Ecco un esempio:

```html
<audio id="meowClip" controls>
  <source src="audio/meow.mp3" type="audio/mpeg">
  <source src="audio/meow.ogg" type="audio/ogg">
</audio>
```

**Nota:** I contenuti multimediali di solito hanno componenti sia visive sia uditive. Hanno bisogno di sottotitoli sincronizzati e di una trascrizione affinché gli utenti con disabilità visive e/o uditive possano consultarli. Generalmente, lo sviluppatore web non è responsabile della creazione dei sottotitoli o della trascrizione, ma deve sapere di includerli.

# --instructions--

È ora di fare una pausa da Camper Cat e incontrare il compagno camper Zersiax (@zersiax), un campione di accessibilità e utente di lettore di schermo. Per ascoltare un estratto del suo lettore di schermo in azione, aggiungi un elemento `audio` dopo l'elemento `p`. Includi l'attributo `controls`. Poi inserisci un elemento `source` all'interno dei tag `audio` con l'attributo `src` impostato su `https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3` e l'attributo `type` impostato su `"audio/mpeg"`.

**Nota:** Il clip audio potrebbe sembrare veloce e difficile da capire, ma questa è una velocità normale per gli utenti di lettori di schermo.

# --hints--

Il tuo codice dovrebbe avere un tag `audio`.

```js
assert.lengthOf(document.querySelectorAll('audio'),1);
```

Il tuo elemento `audio` dovrebbe avere un tag di chiusura.

```js
assert.match(code,/<audio.*>[\s\S]*<\/audio>/g);
assert.lengthOf(code.match(/<\/audio>/g),1);
```

Il tag `audio` dovrebbe avere l'attributo `controls`.

```js
assert.exists(document.querySelector('audio')?.getAttribute('controls'));
```

Il tuo codice dovrebbe avere un tag `source`.

```js
assert.lengthOf(document.querySelectorAll('source'), 1);
```

Il tuo tag `source` dovrebbe essere all'interno dei tag `audio`.

```js
const audio = document.querySelector('audio');
const children = audio.querySelectorAll(`:scope ${'source'}`);
assert.lengthOf(children,1);
```

Il valore per l'attributo `src` sul tag `source` dovrebbe corrispondere esattamente al collegamento nelle istruzioni.

```js
assert.equal(
  document.querySelector('source')?.getAttribute('src'),
  'https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3'
);
```

Il tuo codice dovrebbe includere un attributo `type` sul tag `source` con valore audio/mpeg.

```js
assert.equal(document.querySelector('source')?.getAttribute('type'), 'audio/mpeg');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>



  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Real Coding Ninjas</h1>
  </header>
  <main>
    <p>A sound clip of Zersiax's screen reader in action.</p>
    <audio controls>
      <source src="https://cdn.freecodecamp.org/curriculum/applied-accessibility/screen-reader.mp3" type="audio/mpeg"/>
    </audio>
  </main>
</body>
```
