---
id: 587d778a367417b2b2512aa6
title: "migliorare l'accessibilità dei campi modulo con l'elemento label"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMMAN'
forumTopicId: 301016
dashedName: improve-form-field-accessibility-with-the-label-element
---

# --description--

Migliorare l'accessibilità con il markup HTML semantico riguarda sia l'uso di nomi di tag appropriati sia di attributi. Le prossime sfide coprono alcuni scenari importanti sull'uso degli attributi nei moduli.

Il tag `label` avvolge il testo per un controllo specifico del modulo, di solito il nome o l'etichetta di una scelta. Questo lega un significato all'elemento e rende il modulo più leggibile. L'attributo `for` su un tag `label` associa esplicitamente quel `label` al controllo del modulo ed è usato dai lettori di schermo.

Hai imparato a conoscere i pulsanti di opzione e le loro etichette in una lezione nella sezione Basic HTML. In quella lezione, abbiamo avvolto l'elemento input del pulsante di opzione dentro un elemento `label` insieme al testo dell'etichetta per rendere il testo cliccabile. Un altro modo per ottenere questo è usando l'attributo `for`, come spiegato in questa lezione.

Il valore dell'attributo `for` deve essere lo stesso del valore dell'attributo `id` del controllo del modulo. Ecco un esempio:

```html
<form>
  <label for="name">Name:</label>
  <input type="text" id="name" name="name">
</form>
```

# --instructions--

Camper Cat si aspetta molto interesse per i suoi post riflessivi sul blog e vuole includere un modulo di iscrizione via email. Aggiungi un attributo `for` sull'`label` email che corrisponda all'`id` sul suo campo `input`.

# --hints--

Il tuo codice dovrebbe avere un attributo `for` sul tag `label` che non sia vuoto.

```js
assert.isNotEmpty(document.querySelector('label')?.getAttribute('for'));
```

Il valore del tuo attributo `for` dovrebbe corrispondere al valore `id` sull'input email `input`.

```js
assert.equal(document.querySelector('label')?.getAttribute('for'), 'email');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label>Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>


      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
