---
id: 697fc395f3d85f14b05cd9ca
title: Restaurar una narrativa coherente a partir de un arreglo de fragmentos de historia
challengeType: 26
dashedName: lab-story-fragment-restoration
---

# --description--

En este laboratorio, restaurarás una narrativa coherente a partir de un arreglo corrupto de fragmentos de historia.

Practicarás el uso de ciclos implementando algoritmos fundamentales para arreglos desde cero.

Trabajarás con arreglos de objetos fragmento de historia. Cada objeto fragmento tiene las siguientes propiedades:

| Propiedad | Descripción | Valor de ejemplo |
|-|-|-|
| `id` | Un entero positivo que indica la posición del fragmento en la historia | `3` |
| `text` | El contenido real de la historia | `"and I use Arch btw.\""` |

En este laboratorio, se te proporciona un arreglo prellenado llamado `shuffledFragments`.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de usuario**

1. No debes cambiar el arreglo prellenado `shuffledFragments`.
2. Debes crear una función llamada `compactFragments` que reciba un arreglo de fragmentos y devuelva un nuevo arreglo con todos los elementos undefined eliminados. Si la función elimina algún elemento undefined, debe registrar un mensaje en la consola. El mensaje debe comenzar con el prefijo `[COMPACTED]`.
3. Debes declarar una variable llamada `compactedShuffledFragments` y asignarle el resultado de llamar a `compactFragments` con el arreglo `shuffledFragments`.
4. Debes crear una función llamada `sortFragments` que reciba un arreglo de fragmentos sin elementos `undefined` y devuelva un nuevo arreglo ordenado por la propiedad `id` de forma ascendente, manteniendo los fragmentos que comparten el mismo `id` en su orden original. No debes usar el método `sort` integrado de JavaScript.
5. Debes declarar una variable llamada `sortedFragments` y asignarle el resultado de llamar a `sortFragments` con el arreglo `compactedShuffledFragments`.
6. Debes crear una función llamada `dedupeFragments` que reciba un arreglo ordenado de fragmentos y devuelva un nuevo arreglo con los duplicados eliminados, conservando solo la primera aparición. Debes definir los duplicados como dos o más fragmentos que comparten el mismo `id`. Por cada `id` que se deduplique, la función debe registrar un mensaje en la consola. El mensaje debe comenzar con el prefijo `[DEDUPED]`.
7. Debes declarar una variable llamada `dedupedFragments` y asignarle el resultado de llamar a `dedupeFragments` con el arreglo `sortedFragments`.
8. Debes crear una función llamada `fillMissingFragments` que reciba un arreglo ordenado de fragmentos y devuelva un nuevo arreglo con los fragmentos faltantes rellenados con objetos de marcador de posición. Debes definir los fragmentos faltantes como huecos en la secuencia entre el más bajo y el más alto `id`. Los objetos de marcador de posición deben tener el formato `{ id: missingId, text: "[...]" }`. Por cada marcador de posición añadido, la función debe registrar un mensaje en la consola. El mensaje debe comenzar con el prefijo `[FILLED]`.
9. Debes declarar una variable llamada `filledFragments` y asignarle el resultado de llamar a `fillMissingFragments` con el arreglo `dedupedFragments`.
10. Debes crear una función llamada `assembleStory` que reciba un arreglo ordenado de fragmentos y devuelva una única cadena que contenga todos los textos de los fragmentos, separados por saltos de línea.
11. Debes usar `assembleStory` con tus `filledFragments` para mostrar la historia completa en la consola.
12. Tus funciones `compactFragments`, `sortFragments`, `dedupeFragments`, `fillMissingFragments` y `assembleStory` no deben mutar el arreglo con el que se les llama.

**Ejemplo**

Aquí tienes un ejemplo de un arreglo que contiene fragmentos de historia:

```js
const exampleArray = [
  { id: 3, text: "and I use Arch btw.\"" },
  ,
  { id: 1, text: "Naomi said:" },
  { id: 3, text: "and I use Arch btw.\"" },
];
```

Después de restaurar la historia desde `exampleArray`, se vería así:

```md
Naomi said:
[...]
and I use Arch btw."
```

# --hints--

No debes cambiar el arreglo prellenado `shuffledFragments`.

```js
const shuffledFragments__unchanged = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

assert.deepEqual(shuffledFragments, shuffledFragments__unchanged);
```

Debes tener una función `compactFragments`.

```js
assert.isFunction(compactFragments);
```

`compactFragments([{ id: 1, text: "Hello" }, undefined])` debe registrar un mensaje que comience con `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Cuando la función `compactFragments` se llama con un arreglo que contiene elementos `undefined`, debe registrar un mensaje en la consola que comience con `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" }, undefined, { id: 11, text: "bar" }, undefined]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" },,])` debe registrar un mensaje que comience con `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 1, text: "Hello" },,]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Cuando la función `compactFragments` se llama con un arreglo que contiene espacios vacíos, debe registrar un mensaje en la consola que comience con `[COMPACTED]`.

```js
const spy = __helpers.spyOn(console, 'log');
compactFragments([{ id: 10, text: "foo" },,{ id: 11, text: "bar" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

Cuando la función `compactFragments` se llama con el arreglo `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`, no debe registrar ningún mensaje en la consola que comience con `[COMPACTED]`.

```js
const shuffledFragments__withoutUndefined = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const spy = __helpers.spyOn(console, 'log');
compactFragments(shuffledFragments__withoutUndefined);
assert.isFalse(spy.calls.some(call => call[0].match(/^\[COMPACTED\]/)));
spy.restore();
```

`compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }])` debe devolver `[{ id: 1, text: "Hello" }, { id: 2, text: "World" }]`.

```js
assert.deepEqual(
  compactFragments([{ id: 1, text: "Hello" }, undefined, { id: 2, text: "World" }]),
  [{ id: 1, text: "Hello" }, { id: 2, text: "World" }]
);
```

`compactFragments` no debe mutar el arreglo que recibe.

```js
const input = [{ id: 1, text: "a" }, undefined, { id: 2, text: "b" }];
const copy = structuredClone(input);
compactFragments(input);
assert.deepEqual(input, copy);
```

`compactedShuffledFragments` debe ser un arreglo sin elementos `undefined`.

```js
assert.isArray(compactedShuffledFragments);
assert.isFalse(compactedShuffledFragments.some(f => f === undefined));
```

Debes tener una función `sortFragments`.

```js
assert.isFunction(sortFragments);
```

`sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }])` debe devolver `[{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]`.

```js
assert.deepEqual(
  sortFragments([{ id: 30, text: "c" }, { id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }]),
  [{ id: 10, text: "a" }, { id: 10, text: "dup" }, { id: 20, text: "b" }, { id: 30, text: "c" }]
);
```

`sortFragments` no debe usar el método integrado `.sort()`.

```js
function isBuiltInSortUsed() {
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => { sortUsed = true; };
  try {
    sortFragments([{ id: 2, text: "b" }, { id: 1, text: "a" }]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

`sortFragments` no debe mutar el arreglo que recibe.

```js
const input = [{ id: 3, text: "c" }, { id: 1, text: "a" }];
const copy = structuredClone(input);
sortFragments(input);
assert.deepEqual(input, copy);
```

`sortedFragments` debe tener cada `id` de elemento menor o igual al `id` del siguiente elemento.

```js
for (let i = 0; i < sortedFragments.length - 1; i++) {
  assert.isAtMost(sortedFragments[i].id, sortedFragments[i + 1].id);
}
```

Debes tener una función `dedupeFragments`.

```js
assert.isFunction(dedupeFragments);
```

`dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }])` debe devolver `[{ id: 1, text: "first" }, { id: 2, text: "second" }]`.

```js
assert.deepEqual(
  dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]),
  [{ id: 1, text: "first" }, { id: 2, text: "second" }]
);
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }])` debe registrar un mensaje `[DEDUPED]` por cada `id` duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }, { id: 2, text: "dup" }, { id: 3, text: "c" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 2);
spy.restore();
```

`dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }])` debe registrar un mensaje `[DEDUPED]` cuando el último `id` está duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "a" }, { id: 2, text: "b" }, { id: 2, text: "dup" }]);
const dedupedLogs = spy.calls.filter(call => call[0].match(/^\[DEDUPED\]/));
assert.equal(dedupedLogs.length, 1);
spy.restore();
```

Cuando `dedupeFragments` se llama con un arreglo que contiene duplicados, debe registrar un mensaje que comience con `[DEDUPED]` por cada `id` duplicado.

```js
const spy = __helpers.spyOn(console, 'log');
dedupeFragments([{ id: 1, text: "first" }, { id: 1, text: "dup" }, { id: 2, text: "second" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[DEDUPED\]/)));
spy.restore();
```

`dedupeFragments` no debe mutar el arreglo que recibe.

```js
const input = [{ id: 1, text: "a" }, { id: 1, text: "dup" }, { id: 2, text: "b" }];
const copy = structuredClone(input);
dedupeFragments(input);
assert.deepEqual(input, copy);
```

`dedupedFragments` no debe tener valores `id` duplicados.

```js
const ids = dedupedFragments.map(f => f.id);
assert.strictEqual(ids.length, new Set(ids).size);
```

Debes tener una función `fillMissingFragments`.

```js
assert.isFunction(fillMissingFragments);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` debe devolver `[{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]`.

```js
assert.deepEqual(
  fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]),
  [{ id: 1, text: "a" }, { id: 2, text: "[...]" }, { id: 3, text: "c" }]
);
```

`fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }])` debe registrar un mensaje que comience con `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 1, text: "a" }, { id: 3, text: "c" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

Cuando `fillMissingFragments` agrega un marcador de posición, debe registrar un mensaje que comience con `[FILLED]`.

```js
const spy = __helpers.spyOn(console, 'log');
fillMissingFragments([{ id: 5, text: "e" }, { id: 8, text: "h" }]);
assert.isTrue(spy.calls.some(call => call[0].match(/^\[FILLED\]/)));
spy.restore();
```

`fillMissingFragments` no debe mutar el arreglo que recibe.

```js
const input = [{ id: 1, text: "a" }, { id: 3, text: "c" }];
const copy = structuredClone(input);
fillMissingFragments(input);
assert.deepEqual(input, copy);
```

`filledFragments` no debe tener espacios en la secuencia de `id`.

```js
for (let i = 0; i < filledFragments.length - 1; i++) {
  assert.strictEqual(filledFragments[i + 1].id, filledFragments[i].id + 1);
}
```

Debes tener una función `assembleStory`.

```js
assert.isFunction(assembleStory);
```

`assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }])` debe devolver la cadena `"Hello\nWorld"`.

```js
assert.strictEqual(
  assembleStory([{ id: 1, text: "Hello" }, { id: 2, text: "World" }]),
  "Hello\nWorld"
);
```

`assembleStory` no debe mutar el arreglo que recibe.

```js
const input = [{ id: 1, text: "Hello" }, { id: 2, text: "World" }];
const copy = structuredClone(input);
assembleStory(input);
assert.deepEqual(input, copy);
```

Debes llamar a `console.log` con el resultado de `assembleStory(filledFragments)`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(.*assembleStory\s*\(\s*filledFragments\s*\)/
);
```

# --seed--

## --seed-contents--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];
```

# --solutions--

```js
const shuffledFragments = [
  { id: 15, text: "and, after a time, passed the place where the Hare was sleeping." },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 11, text: "and to make the Tortoise feel very deeply how ridiculous it was for him to try a race with a Hare," },
  { id: 7, text: "but for the fun of the thing he agreed." },
  { id: 19, text: "The Hare now ran his swiftest," },
  ,
  { id: 1, text: "A Hare was making fun of the Tortoise one day for being so slow." },
  { id: 14, text: "The Tortoise meanwhile kept going slowly but steadily," },
  { id: 9, text: "marked the distance and started the runners off." },
  ,
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 17, text: "and when at last he did wake up," },
  { id: 2, text: '"Do you ever get anywhere?" he asked with a mocking laugh.' },
  { id: 12, text: "he lay down beside the course to take a nap" },
  ,
  { id: 8, text: "So the Fox, who had consented to act as judge," },
  { id: 20, text: "but he could not overtake the Tortoise in time." },
  { id: 5, text: "I'll run you a race and prove it.\"" },
  { id: 6, text: "The Hare was much amused at the idea of running a race with the Tortoise," },
  ,
  { id: 13, text: "until the Tortoise should catch up." },
  { id: 10, text: "The Hare was soon far out of sight," },
  { id: 12, text: "he lay down beside the course to take a nap" },
  { id: 18, text: "the Tortoise was near the goal." },
];

const compactFragments = (fragments) => {
  let compacted = [];
  let numberOfUndefined = 0;

  for (const fragment of fragments) {
    if (fragment === undefined) {
      numberOfUndefined++;
    } else {
      compacted.push(fragment);
    }
  }

  if (numberOfUndefined > 0) {
    console.log(`[COMPACTED] Removed ${numberOfUndefined} undefined fragment(s).`);
  }

  return compacted;
};

const compactedShuffledFragments = compactFragments(shuffledFragments);

const sortFragments = (fragments) => {
  let sorted = [...fragments];

  for (let i = 0; i < sorted.length; i++) {
    for (let j = i; j > 0; j--) {
      if (sorted[j - 1].id > sorted[j].id) {
        swap(sorted, j - 1, j);
      }
    }
  }

  function swap(array, index1, index2) {
    [sorted[index1], sorted[index2]] = [sorted[index2], sorted[index1]];
  }

  return sorted;
};

const sortedFragments = sortFragments(compactedShuffledFragments);

const dedupeFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let deduped = [shallowCopy.shift()];
  let duplicates = [];
  let currentIdOccurrence = 1;

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = deduped.at(-1);
    if (fragmentToInsert.id === previousFragment.id) {
      currentIdOccurrence++;
    } else {
      deduped.push(fragmentToInsert);

      if (currentIdOccurrence > 1) {
        duplicates.push({ id: previousFragment.id, occurrence: currentIdOccurrence });
      }

      currentIdOccurrence = 1;
    }
  }

  if (currentIdOccurrence > 1) {
    duplicates.push({ id: deduped.at(-1).id, occurrence: currentIdOccurrence });
  }

  for (const duplicate of duplicates) {
    console.log(`[DEDUPED] Fragment #${duplicate.id} appeared ${duplicate.occurrence} times, kept first occurrence.`);
  }

  return deduped;
};

const dedupedFragments = dedupeFragments(sortedFragments);

const fillMissingFragments = (sortedFragments) => {
  const shallowCopy = [...sortedFragments];
  let filled = [shallowCopy.shift()];

  for (const fragmentToInsert of shallowCopy) {
    const previousFragment = filled.at(-1);
    const numberOfMissingIDs = fragmentToInsert.id - 1 - previousFragment.id;
    for (let i = 1; i <= numberOfMissingIDs; i++) {
      const missingId = previousFragment.id + i;
      filled.push({ id: missingId, text: "[...]" });
      console.log(`[FILLED] Fragment #${missingId} missing, placeholder inserted.`)
    }
    filled.push(fragmentToInsert);
  }

  return filled;
};

const filledFragments = fillMissingFragments(dedupedFragments);

const assembleStory = (sortedFragments) => {
  const texts = [];
  for (const fragment of sortedFragments) {
    texts.push(fragment.text);
  }
  return texts.join("\n");
};

console.log("\nHere is the assembled story:\n");
console.log(assembleStory(filledFragments));
```
