---
id: 680b3ef395479b0e449ecb6e
title: Implementieren Sie den Selection-Sort-Algorithmus
challengeType: 27
dashedName: implement-selection-sort-algorithm
---

# --description--

Selection Sort ist ein weiterer beliebter Sortieralgorithmus, der in den meisten Informatikkursen gelehrt wird.

Dieser Algorithmus funktioniert, indem er wiederholt das kleinste Element aus dem unsortierten Teil der Liste findet und es mit dem ersten unsortierten Element vertauscht. Er beginnt damit, den Minimalwert in der gesamten Liste auszuwählen und mit dem ersten Element zu tauschen. Dann geht er zur zweiten Position, findet den kleinsten Wert in den verbleibenden unsortierten Elementen und tauscht ihn mit dem zweiten Element. Dieser Vorgang wird fortgesetzt, indem die Liste Element für Element durchlaufen wird, bis die gesamte Liste sortiert ist.

Selection Sort hat eine quadratische Zeitkomplexität im besten, durchschnittlichen und schlechtesten Fall. Die Speicherkomplexität ist konstant `O(1)`, da die Sortierung in-place erfolgt und unabhängig von der Größe der Liste nur eine konstante Menge an Speicher verwendet wird.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um die Aufgabe abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `selection_sort` definieren.
1. Ihre `selection_sort`-Funktion sollte einen Parameter haben, der die Liste der Elemente darstellt.
1. Ihre `selection_sort`-Funktion sollte eine Liste entgegennehmen und die Elemente in-place von klein nach groß sortieren.
1. Ihre `selection_sort`-Funktion sollte die Eingabeliste in-place ändern und sie zurückgeben, sobald sie sortiert ist.
1. Ihre `selection_sort`-Funktion sollte dem Selection-Sort-Algorithmus folgen, indem sie das kleinste Element aus dem unsortierten Teil der Liste mit dem ersten unsortierten Element tauscht.
1. Ihre `selection_sort`-Funktion sollte keine unnötigen Vertauschungen durchführen, wenn das kleinste Element bereits an der richtigen Position ist.
1. Ihre `selection_sort`-Funktion sollte weder die eingebaute `sort()`-Methode noch die `sorted()`-Funktion verwenden.

# --hints--

Sie sollten eine Funktion mit dem Namen `selection_sort` haben.

```js
({test: () => {
    runPython(`
assert _Node(_code).has_function('selection_sort')
    `)
}})
```

Ihre `selection_sort`-Funktion sollte einen Parameter haben.

```js
(
    {test: () => {
        runPython(`
            import inspect
            sig = inspect.signature(selection_sort)
            assert len(sig.parameters) == 1
        `)
    }}
)
```

Sie sollten kein Modul importieren oder eingebaute Sortiermethoden in Ihrem Code verwenden.

```js
(
    {
        test: () => runPython(`
            assert len(_Node(_code).find_imports()) == 0
            assert not _Node(_code).block_has_call("sort")
            assert not _Node(_code).block_has_call("sorted")
        `)
    }
)
```

Ihre `selection_sort` sollte dieselbe Liste zurückgeben wie die Eingabeliste.

```js
(
    {
        test: () => runPython(`
            input_list = [32, 0, 88, 1, 203]
            assert selection_sort(input_list) is input_list
        `)
    }
)
```

Ihre `selection_sort` sollte die Eingabeliste in-place ändern. Sie sollten keine Methode verwenden, die Elemente zur Liste hinzufügt oder entfernt.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)

                def append(self, *args, **kwargs):
                    self._record.append(f'Appending to list: {args}, {kwargs}')
                    super().append(*args, **kwargs)

                def insert(self, *args, **kwargs):
                    self._record.append(f'Inserting to list: {args}, {kwargs}')
                    super().insert(*args, **kwargs)

                def pop(self, *args, **kwargs):
                    self._record.append(f'Popping from list: {args}, {kwargs}')
                    super().pop(*args, **kwargs)

                def remove(self, *args, **kwargs):
                    self._record.append(f'Removing from list: {args}, {kwargs}')
                    super().remove(*args, **kwargs)


            list_to_sort = CustomList([32, 0, 88, 1, 203])

            selection_sort(list_to_sort)
            assert all(
                'Setting item' in record
                for record in list_to_sort._record
            )
        `)
    }
)
```

Ihre `selection_sort`-Funktion sollte dem Selection-Sort-Algorithmus folgen, indem sie den Minimalwert im unsortierten Teil der Liste mit dem ersten unsortierten Element tauscht. Vermeiden Sie unnötige Vertauschungen, wenn der Minimalwert bereits an der richtigen Position ist.

```js
(
    {
        test: () => runPython(`
            from collections import UserList

            class CustomList(UserList):
                def __init__(self, *args, **kwargs):
                    self._record = []
                    super().__init__(*args, **kwargs)

                def __setitem__(self, key, value):
                    self._record.append(f'Setting item {key} to {value}. Previous value: {self[key]}')
                    super().__setitem__(key, value)


            list_to_sort = CustomList([33, 1, 89, 2, 67, 245])

            swap_pairs = [
                ((0, 1, 33), (1, 33, 1)),
                ((1, 2, 33), (3, 33, 2)),
                ((2, 33, 89), (3, 89, 33)),
                ((3, 67, 89), (4, 89, 67)),
            ]

            selection_sort(list_to_sort)
            record = list_to_sort._record

            actual_pairs = [
                (record[offset * 2], record[offset * 2 + 1])
                for offset, _ in enumerate(record[::2])
            ]
            assert len(swap_pairs) == len(actual_pairs)

            for expected_pair, actual_pair in zip(swap_pairs, actual_pairs):
                assert (
                    set(f'Setting item {key} to {value}. Previous value: {old}' for key, value, old in expected_pair)
                    == set(actual_pair)
                )
        `)
    }
)
```

`selection_sort([33, 1, 89, 2, 67, 245])` sollte `[1, 2, 33, 67, 89, 245]` zurückgeben.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([33, 1, 89, 2, 67, 245]) == [1, 2, 33, 67, 89, 245]
            
        `)
    }
)
```

`selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3])` sollte `[3, 5, 12, 15, 16, 23, 72, 99, 567]` zurückgeben.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([5, 16, 99, 12, 567, 23, 15, 72, 3]) == [3, 5, 12, 15, 16, 23, 72, 99, 567]
            
        `)
    }
)
```

`selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` sollte `[1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]` zurückgeben.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]) == [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
            
        `)
    }
)
```

Ihre `selection_sort`-Funktion sollte jede Liste von Zahlen korrekt sortieren.

```js
(
    {
        test: () => runPython(`
            assert selection_sort([42, 17, 93, 8, 61, 29]) == [8, 17, 29, 42, 61, 93]
            assert selection_sort([11, 4, 78, 23, 55, 198, 65, 90, 2]) == [2, 4, 11, 23, 55, 65, 78, 90, 198]
            assert selection_sort([9, 27, 3, 7, 101, 66, 34, 52, 87, 42, 12, 29]) == [3, 7, 9, 12, 27, 29, 34, 42, 52, 66, 87, 101]
            assert selection_sort([5, 14, 33, 77, 2, 18, 92, 1, 100, 45, 73, 64, 28, 56]) == [1, 2, 5, 14, 18, 28, 33, 45, 56, 64, 73, 77, 92, 100]
            
        `)
    }
)     
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def selection_sort(nums):
    for i, _ in enumerate(nums):
        min_index = i

        for j in range(i + 1, len(nums)):
            if nums[j] < nums[min_index]:
                min_index = j

        if min_index != i:
            nums[i], nums[min_index] = nums[min_index], nums[i]

    return nums

```
