---
id: a39963a4c10bc8b4d4f06d7e
title: 实现一个值移除函数
challengeType: 26
dashedName: implement-a-value-remover-function
---

# --description--

在本实验中，你将创建一个函数，该函数以初始数组作为第一个参数，后跟一个或多个表示要移除值的附加参数。

**目标**：完成以下用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个 `destroyer` 函数，它接受一个数组和一个或多个额外的参数。
1. `destroyer` 函数应返回一个新建数组，排除第一个参数中与后续任何参数匹配的所有元素。
1. 该函数必须接受不确定数量的参数。

# --hints--

你应该有一个 `destroyer` 函数。

```js
assert.isFunction(destroyer);
```

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)` 应返回 `[1, 1]`。

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)` 应返回 `[1, 5, 1]`。

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)` 应返回 `[1]`。

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)` 应返回 `[]`。

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)` 应返回 `["hamburger"]`。

```js
assert.deepEqual(destroyer(["tree", "hamburger", 53], "tree", 53), ["hamburger"]);
```

`destroyer( ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan" )` 应该返回 `[12, 92, 65]`。

```js
assert.deepEqual(
  destroyer(
    ["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"],
    "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan"
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function destroyer(arr, ...valsToRemove) {
  return arr.filter(item => !valsToRemove.includes(item));
}
```

