---
id: 68dbc6b7e8f794a2ae69bc74
title: Nクイーンアルゴリズムの実装
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

Nクイーン問題は、N×NのチェスボードにN個のクイーンを配置して、どのクイーンも互いに攻撃し合わないようにする（同じ行、列、斜めに2つのクイーンが存在しない）問題です。

例えば、4×4のボードがある場合、1つの有効な配置は次の通りです。

```md
[1, 3, 0, 2]
```

これは、行0のクイーンが列1に、行1のクイーンが列3に、行2のクイーンが列0に、行3のクイーンが列2に配置されていることを意味します。

視覚的には、この配置は次のように見えます。

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

ここで、`Q` はクイーンを、`.` は空のマスを表します。

このラボでは、深さ優先探索のアプローチを使ってNクイーン問題の解決プログラムを実装します。

**目的**：以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `dfs_n_queens` という名前の関数が必要です。
2. その関数は引数を1つだけ受け取る必要があります：整数の `n` です。
3. `n` が `1` より小さい場合、関数は空のリスト（`[]`）を返す必要があります。
4. 関数は解のリストを返す必要があります。各解は長さ `n` のリストで、インデックス `i` の要素は行 `i` にあるクイーンの列インデックス（0始まり）を表します。

# --hints--

`dfs_n_queens` という名前の関数が1つの引数を受け取る必要があります。

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

`n` が `1` より小さい場合、関数は空のリストを返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

関数は解のリストを返す必要があり、各解は長さ `n` のリストです。

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` は `[[0]]` を返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` は `[]` を返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` は `[]` を返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` は `[[1, 3, 0, 2], [2, 0, 3, 1]]` を返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` は `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]` を返す必要があります。

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` は `10` である必要があります。

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` は `92` である必要があります。

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` は正しい結果を返す必要があります。

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
