---
id: a5229172f011153519423690
title: Costruisci un calcolatore della somma dei numeri dispari di Fibonacci
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

In questo laboratorio costruirai un calcolatore della somma dei numeri dispari di Fibonacci che calcola la somma di tutti i numeri dispari di Fibonacci minori o uguali a un dato numero intero positivo.

**Obiettivo**: Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente**

1. Dovresti avere una funzione `sumFibs` che accetta un numero come argomento.
1. La funzione `sumFibs` dovrebbe restituire la somma di tutti i numeri dispari di Fibonacci minori o uguali al numero dato.
1. La sequenza di Fibonacci inizia con `0` e `1`, e ogni numero successivo è la somma dei due precedenti.
1. Solo i numeri dispari di Fibonacci dovrebbero essere aggiunti alla somma.

# --hints--

Dovresti avere una funzione `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` dovrebbe restituire un numero.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` dovrebbe restituire `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` dovrebbe restituire `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` dovrebbe restituire `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` dovrebbe restituire `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` dovrebbe restituire `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
