---
id: a6b0bb188d873cb2c8729495
title: HTMLエンティティコンバーターを実装する
challengeType: 26
dashedName: implement-an-html-entity-converter
---

# --description--

このラボでは、文字列内の特殊文字を対応するHTMLエンティティに変換することについて学びます。

**目的**：以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ストーリー：**

1. `convertHTML` 関数が文字列を引数として受け取る必要があります。
1. `convertHTML` 関数は、引数の文字列内の特殊文字を対応するHTMLエンティティに変換した新しい文字列を返す必要があります。

   - `&` は `&amp;` に変換されるべきです。
   - `<` は `&lt;` に変換されるべきです。
   - `>` は `&gt;` に変換されるべきです。
   - `"` は `&quot;` に変換されるべきです。
   - `'` は `&apos;` に変換されるべきです。

# --hints--

`convertHTML` 関数を用意してください。

```js
assert.isFunction(convertHTML);
```

`convertHTML("Dolce & Gabbana")` は文字列 `Dolce &amp; Gabbana` を返すべきです。

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")` は文字列 `Hamburgers &lt; Pizza &lt; Tacos` を返すべきです。

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")` は文字列 `Sixty &gt; twelve` を返すべきです。

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')` は文字列 `Stuff in &quot;quotation marks&quot;` を返すべきです。

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")` は文字列 `Schindler&apos;s List` を返すべきです。

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")` は文字列 `&lt;&gt;` を返すべきです。

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")` は文字列 `abc` を返すべきです。

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const map = {
  '&': '&amp;',
  '<': '&lt;',
  '>': '&gt;',
  '"': '&quot;',
  "'": '&apos;'
};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(char) {
    return map[char];
  });
}
```
