---
id: a103376db3ba46b2d50db289
title: 스파이널 탭 케이스
challengeType: 1
forumTopicId: 16078
dashedName: spinal-tap-case
---

# --description--

문자열을 스파이널 케이스로 변환하세요. 스파이널 케이스는 모두 소문자로 된 단어들이 대시로 연결된 형태입니다.

# --hints--

`spinalCase("This Is Spinal Tap")`는 문자열 `this-is-spinal-tap`을 반환해야 합니다.

```js
assert.deepEqual(spinalCase('This Is Spinal Tap'), 'this-is-spinal-tap');
```

`spinalCase("thisIsSpinalTap")`는 문자열 `this-is-spinal-tap`을 반환해야 합니다.

```js
assert.strictEqual(spinalCase('thisIsSpinalTap'), 'this-is-spinal-tap');
```

`spinalCase("The_Andy_Griffith_Show")`는 문자열 `the-andy-griffith-show`을 반환해야 합니다.

```js
assert.strictEqual(
  spinalCase('The_Andy_Griffith_Show'),
  'the-andy-griffith-show'
);
```

`spinalCase("Teletubbies say Eh-oh")`는 문자열 `teletubbies-say-eh-oh`을 반환해야 합니다.

```js
assert.strictEqual(
  spinalCase('Teletubbies say Eh-oh'),
  'teletubbies-say-eh-oh'
);
```

`spinalCase("AllThe-small Things")`는 문자열 `all-the-small-things`을 반환해야 합니다.

```js
assert.strictEqual(spinalCase('AllThe-small Things'), 'all-the-small-things');
```

# --seed--

## --seed-contents--

```js
function spinalCase(str) {
  return str;
}

spinalCase('This Is Spinal Tap');
```

# --solutions--

```js
function spinalCase(str) {
  str = str.replace(/([a-z](?=[A-Z]))/g, '$1 ');
  return str.toLowerCase().replace(/\ |\_/g, '-');
}
```
