---
id: a105e963526e7de52b219be9
title: 정렬된 합집합
challengeType: 1
forumTopicId: 16077
dashedName: sorted-union
---

# --description--

두 개 이상의 배열을 받아 원래 제공된 배열들의 순서대로 고유한 값들만 포함하는 새 배열을 반환하는 함수를 작성하세요.

즉, 모든 배열에 존재하는 모든 값이 원래 순서대로 포함되어야 하지만, 최종 배열에는 중복이 없어야 합니다.

고유한 숫자들은 원래 순서대로 정렬되어야 하지만, 최종 배열은 숫자 순서로 정렬되지 않아야 합니다.

예시는 어서션 테스트를 확인하세요.

# --hints--

`uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1])`는 `[1, 3, 2, 5, 4]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

`uniteUnique([1, 2, 3], [5, 2, 1])`는 `[1, 2, 3, 5]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1]), [1, 2, 3, 5]);
```

`uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8])`는 `[1, 2, 3, 5, 4, 6, 7, 8]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 2, 3], [5, 2, 1, 4], [2, 1], [6, 7, 8]), [
  1,
  2,
  3,
  5,
  4,
  6,
  7,
  8
]);
```

`uniteUnique([1, 3, 2], [5, 4], [5, 6])`는 `[1, 3, 2, 5, 4, 6]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2], [5, 4], [5, 6]), [1, 3, 2, 5, 4, 6]);
```

`uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1])`는 `[1, 3, 2, 5, 4]`를 반환해야 합니다.

```js
assert.deepEqual(uniteUnique([1, 3, 2, 3], [5, 2, 1, 4], [2, 1]), [1, 3, 2, 5, 4]);
```

# --seed--

## --seed-contents--

```js
function uniteUnique(arr) {
  return arr;
}

uniteUnique([1, 3, 2], [5, 2, 1, 4], [2, 1]);
```

# --solutions--

```js
function uniteUnique(arr) {
  return [].slice.call(arguments).reduce(function(a, b) {
    return [].concat(
      a, 
      b.filter(function(e, currentIndex) {
        return b.indexOf(e) === currentIndex && a.indexOf(e) === -1;
      }));
  }, []);
}
```
