---
id: a2f1d72d9b908d0bd72bb9f6
title: 사람 만들기
challengeType: 1
forumTopicId: 16020
dashedName: make-a-person
---

# --description--

아래 메서드들로 객체 생성자를 채우세요:

```js
getFirstName()
getLastName()
getFullName()
setFirstName(first)
setLastName(last)
setFullName(first, last)
```

각 메서드에 대한 예상 출력 결과를 확인하려면 테스트를 실행하세요. 이 메서드들만 객체와 상호작용할 수 있는 유일한 수단이어야 합니다. 각 테스트는 `Person`로서 새로운 `new Person('Bob', 'Ross')` 인스턴스를 선언합니다.

# --hints--

함수 시그니처는 변경하지 마세요.

```js
assert.match(code, /const\s+Person\s*=\s*function\s*\(\s*first\s*,\s*last\s*\)\s*{/);
```

`first` 매개변수를 재할당하지 마세요.

```js
assert.notMatch(code, /\bfirst\s*=\s*/);
```

`last` 매개변수를 재할당하지 마세요.

```js
assert.notMatch(code, /\blast\s*=\s*/);
```

속성을 추가하지 마세요. `Object.keys(Person).length`는 항상 6을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
_person.setLastName('Curry');
_person.setFullName('John', 'Smith');
assert.lengthOf(Object.keys(_person), 6);
```

`Person` 객체를 인스턴스화할 수 있어야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.instanceOf(_person, Person);
```

`Person` 객체에 `firstName` 속성이 없어야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'firstName');
```

`Person` 객체에 `lastName` 속성이 없어야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.notProperty(_person, 'lastName');
```

`.getFirstName()` 메서드는 문자열 `Bob`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFirstName(), 'Bob');
```

`.getLastName()`는 문자열 `Ross`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getLastName(), 'Ross');
```

`.getFullName()` 메서드는 문자열 `Bob Ross`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
assert.strictEqual(_person.getFullName(), 'Bob Ross');
```

`.getFullName()` 메서드는 `Haskell Ross`을 호출한 후 문자열 `.setFirstName('Haskell')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFirstName('Haskell');
assert.strictEqual(_person.getFullName(), 'Haskell Ross');
```

`.getFullName()` 메서드는 `Bob Curry`을 호출한 후 문자열 `.setLastName('Curry')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setLastName('Curry');
assert.strictEqual(_person.getFullName(), 'Bob Curry');
```

`.getFullName()` 메서드는 `Haskell Curry`을 호출한 후 문자열 `.setFullName('Haskell', 'Curry')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFullName(), 'Haskell Curry');
```

`.getFirstName()` 메서드는 `Haskell`을 호출한 후 문자열 `.setFullName('Haskell', 'Curry')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getFirstName(), 'Haskell');
```

`.getLastName()` 메서드는 `Curry`을 호출한 후 문자열 `.setFullName('Haskell', 'Curry')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Haskell', 'Curry');
assert.strictEqual(_person.getLastName(), 'Curry');
```

`.getFullName()` 메서드는 `Emily Martinez de la Rosa`을 호출한 후 문자열 `.setFullName('Emily Martinez', 'de la Rosa')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFullName(), 'Emily Martinez de la Rosa');
```

`.getFirstName()` 속성은 `Emily Martinez`을 호출한 후 문자열 `.setFullName('Emily Martinez', 'de la Rosa')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getFirstName(), 'Emily Martinez');
```

`.getLastName()` 속성은 `de la Rosa`을 호출한 후 문자열 `.setFullName('Emily Martinez', 'de la Rosa')`을 반환해야 합니다.

```js
const _person = new Person('Bob', 'Ross');
_person.setFullName('Emily Martinez', 'de la Rosa');
assert.strictEqual(_person.getLastName(), 'de la Rosa');
```

# --seed--

## --seed-contents--

```js
const Person = function(first, last) {
  this.getFullName = function() {
    return "";
  };
  return "";
};
```

# --solutions--

```js
const Person = function(first, last) {
  let firstName = first;
  let lastName = last;

  this.getFirstName = function(){
    return firstName;
  };

  this.getLastName = function(){
    return lastName;
  };

  this.getFullName = function(){
    return firstName + " " + lastName;
  };

  this.setFirstName = function(str){
    firstName = str;
  };


  this.setLastName = function(str){
    lastName = str;
  };

  this.setFullName = function(first, last){
    firstName = first;
    lastName = last;
  };
};
```
