---
id: a3566b1109230028080c9345
title: 범위 내 모든 숫자 합계
challengeType: 1
forumTopicId: 16083
dashedName: sum-all-numbers-in-a-range
---

# --description--

두 개의 숫자가 들어 있는 배열을 전달받습니다. 이 두 숫자와 그 사이에 있는 모든 숫자의 합을 반환하세요. 항상 더 작은 숫자가 먼저 오지는 않습니다.

예를 들어, `sumAll([4,1])`는 `10`을 반환해야 합니다. 1과 4 사이(양쪽 포함)의 모든 숫자의 합이 `10`이기 때문입니다.

# --hints--

`sumAll([1, 4])`는 숫자를 반환해야 합니다.

```js
assert(typeof sumAll([1, 4]) === 'number');
```

`sumAll([1, 4])`는 10을 반환해야 합니다.

```js
assert.deepEqual(sumAll([1, 4]), 10);
```

`sumAll([4, 1])`는 10을 반환해야 합니다.

```js
assert.deepEqual(sumAll([4, 1]), 10);
```

`sumAll([5, 10])`는 45를 반환해야 합니다.

```js
assert.deepEqual(sumAll([5, 10]), 45);
```

`sumAll([10, 5])`는 45를 반환해야 합니다.

```js
assert.deepEqual(sumAll([10, 5]), 45);
```

# --seed--

## --seed-contents--

```js
function sumAll(arr) {
  return 1;
}

sumAll([1, 4]);
```

# --solutions--

```js
function sumAll(arr) {
  var sum = 0;
  arr.sort(function(a,b) {return a-b;});
  for (var i = arr[0]; i <= arr[1]; i++) {
    sum += i;
  }
  return sum;
}
```
