---
id: a39963a4c10bc8b4d4f06d7e
title: 찾고 제거하기
challengeType: 1
forumTopicId: 16046
dashedName: seek-and-destroy
---

# --description--

첫 번째 인수로 초기 배열이 `destroyer` 함수에 제공되고, 그 뒤에 하나 이상의 인수가 옵니다. 이 인수들과 같은 값을 가진 초기 배열의 모든 요소를 제거하세요.

이 함수는 가변 인수 함수라고도 하는, 인수 개수가 정해지지 않은 함수를 받아야 합니다. 함수 정의에 나머지 매개변수를 추가하거나 `arguments` 객체를 사용해 추가 인수에 접근할 수 있습니다.

# --hints--

`destroyer([1, 2, 3, 1, 2, 3], 2, 3)`는 `[1, 1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([1, 2, 3, 1, 2, 3], 2, 3), [1, 1]);
```

`destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3)`는 `[1, 5, 1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([1, 2, 3, 5, 1, 2, 3], 2, 3), [1, 5, 1]);
```

`destroyer([3, 5, 1, 2, 2], 2, 3, 5)`는 `[1]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([3, 5, 1, 2, 2], 2, 3, 5), [1]);
```

`destroyer([2, 3, 2, 3], 2, 3)`는 `[]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer([2, 3, 2, 3], 2, 3), []);
```

`destroyer(["tree", "hamburger", 53], "tree", 53)`는 `["hamburger"]`를 반환해야 합니다.

```js
assert.deepEqual(destroyer(['tree', 'hamburger', 53], 'tree', 53), [
  'hamburger'
]);
```

`destroyer(["possum", "trollo", 12, "safari", "hotdog", 92, 65, "grandma", "bugati", "trojan", "yacht"], "yacht", "possum", "trollo", "safari", "hotdog", "grandma", "bugati", "trojan")`는 `[12,92,65]`를 반환해야 합니다.

```js
assert.deepEqual(
  destroyer(
    [
      'possum',
      'trollo',
      12,
      'safari',
      'hotdog',
      92,
      65,
      'grandma',
      'bugati',
      'trojan',
      'yacht'
    ],
    'yacht',
    'possum',
    'trollo',
    'safari',
    'hotdog',
    'grandma',
    'bugati',
    'trojan'
  ),
  [12, 92, 65]
);
```

# --seed--

## --seed-contents--

```js
function destroyer(arr) {
  return arr;
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```

# --solutions--

```js
function destroyer(arr) {
  var hash = Object.create(null);
  [].slice.call(arguments, 1).forEach(function(e) {
    hash[e] = true;
  });
  return arr.filter(function(e) { return !(e in hash);});
}

destroyer([1, 2, 3, 1, 2, 3], 2, 3);
```
