---
id: a3bfc1673c0526e06d3ac698
title: 모든 소수의 합
challengeType: 1
forumTopicId: 16085
dashedName: sum-all-primes
---

# --description--

<dfn>소수</dfn>는 1보다 큰 자연수 중에서 약수가 정확히 2개인 수입니다: 1과 자기 자신입니다. 예를 들어, 2는 1과 2로만 나누어지기 때문에 소수입니다. 반면 4는 1, 2, 4로 나누어지므로 소수가 아닙니다.

`sumPrimes`를 다시 작성하여 num 이하의 모든 소수의 합을 반환하게 하세요.

# --hints--

`sumPrimes(10)`은 숫자를 반환해야 합니다.

```js
assert.deepEqual(typeof sumPrimes(10), 'number');
```

`sumPrimes(10)`는 17을 반환해야 합니다.

```js
assert.deepEqual(sumPrimes(10), 17);
```

`sumPrimes(977)`는 73156을 반환해야 합니다.

```js
assert.deepEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js
function sumPrimes(num) {
  return num;
}

sumPrimes(10);
```

# --solutions--

```js
class PrimeSeive {
  constructor(num) {
    const seive = Array(Math.floor((num - 1) / 2)).fill(true);
    const upper = Math.floor((num - 1) / 2);
    const sqrtUpper = Math.floor((Math.sqrt(num) - 1) / 2);

    for (let i = 0; i <= sqrtUpper; i++) {
      if (seive[i]) {
        // Mark value in seive array
        const prime = 2 * i + 3;
        // Mark all multiples of this number as false (not prime)
        const primeSquaredIndex = 2 * i ** 2 + 6 * i + 3;
        for (let j = primeSquaredIndex; j < upper; j += prime) {
          seive[j] = false;
        }
      }
    }

    this._seive = seive;
  }

  isPrime(num) {
    return num === 2
      ? true
      : num % 2 === 0
        ? false
        : this.isOddPrime(num);
  }

  isOddPrime(num) {
    return this._seive[(num - 3) / 2];
  }
};

function sumPrimes(num) {
  const primeSeive = new PrimeSeive(num);

  let sum = 2;
  for (let i = 3; i <= num; i += 2) {
    if (primeSeive.isOddPrime(i)) sum += i;
  }
  return sum;
}

sumPrimes(10);
```
