---
id: a5229172f011153519423690
title: 모든 홀수 피보나치 수의 합
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

양의 정수 `num`가 주어지면 `num` 이하인 모든 홀수 피보나치 수의 합을 반환하세요.

피보나치 수열의 처음 두 숫자는 0과 1입니다. 수열의 다음 숫자는 항상 이전 두 숫자의 합입니다. 피보나치 수열의 처음 일곱 숫자는 0, 1, 1, 2, 3, 5, 8입니다.

예를 들어, `sumFibs(10)`는 `10` 이하인 모든 홀수 피보나치 수가 1, 1, 3, 5이므로 `10`을 반환해야 합니다.

# --hints--

`sumFibs(1)`는 숫자를 반환해야 합니다.

```js
assert(typeof sumFibs(1) === 'number');
```

`sumFibs(1000)`는 1785를 반환해야 합니다.

```js
assert(sumFibs(1000) === 1785);
```

`sumFibs(4000000)`는 4613732를 반환해야 합니다.

```js
assert(sumFibs(4000000) === 4613732);
```

`sumFibs(4)`는 5를 반환해야 합니다.

```js
assert(sumFibs(4) === 5);
```

`sumFibs(75024)`는 60696을 반환해야 합니다.

```js
assert(sumFibs(75024) === 60696);
```

`sumFibs(75025)`는 135721을 반환해야 합니다.

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
