---
id: a6b0bb188d873cb2c8729495
title: HTML 엔티티 변환
challengeType: 1
forumTopicId: 16007
dashedName: convert-html-entities
---

# --description--

문자열에서 `&`, `<`, `>`, `"` (큰따옴표), 그리고 `'` (작은따옴표) 문자를 해당하는 HTML 엔티티로 변환하세요.

# --hints--

`convertHTML("Dolce & Gabbana")`는 문자열 `Dolce &amp; Gabbana`를 반환해야 합니다.

```js
assert.match(convertHTML('Dolce & Gabbana'), /Dolce &amp; Gabbana/);
```

`convertHTML("Hamburgers < Pizza < Tacos")`는 문자열 `Hamburgers &lt; Pizza &lt; Tacos`를 반환해야 합니다.

```js
assert.match(
  convertHTML('Hamburgers < Pizza < Tacos'),
  /Hamburgers &lt; Pizza &lt; Tacos/
);
```

`convertHTML("Sixty > twelve")`는 문자열 `Sixty &gt; twelve`를 반환해야 합니다.

```js
assert.match(convertHTML('Sixty > twelve'), /Sixty &gt; twelve/);
```

`convertHTML('Stuff in "quotation marks"')`는 문자열 `Stuff in &quot;quotation marks&quot;`를 반환해야 합니다.

```js
assert.match(
  convertHTML('Stuff in "quotation marks"'),
  /Stuff in &quot;quotation marks&quot;/
);
```

`convertHTML("Schindler's List")`는 문자열 `Schindler&apos;s List`를 반환해야 합니다.

```js
assert.match(convertHTML("Schindler's List"), /Schindler&apos;s List/);
```

`convertHTML("<>")`는 문자열 `&lt;&gt;`를 반환해야 합니다.

```js
assert.match(convertHTML('<>'), /&lt;&gt;/);
```

`convertHTML("abc")`는 문자열 `abc`를 반환해야 합니다.

```js
assert.strictEqual(convertHTML('abc'), 'abc');
```

# --seed--

## --seed-contents--

```js
function convertHTML(str) {
  return str;
}

convertHTML("Dolce & Gabbana");
```

# --solutions--

```js
var MAP = { '&': '&amp;',
            '<': '&lt;',
            '>': '&gt;',
            '"': '&quot;',
            "'": '&apos;'};

function convertHTML(str) {
  return str.replace(/[&<>"']/g, function(c) {
    return MAP[c];
  });
}
```
