---
id: a97fd23d9b809dac9921074f
title: 인수가 선택 사항인 함수
challengeType: 1
forumTopicId: 14271
dashedName: arguments-optional
---

# --description--

두 인수를 더하는 함수를 만드세요. 만약 인수가 하나만 주어지면, 하나의 인수를 받아 합을 반환하는 함수를 반환하세요.

예를 들어, `addTogether(2, 3)`는 `5`을 반환해야 하고, `addTogether(2)`는 함수를 반환해야 합니다.

이 반환된 함수를 하나의 인수로 호출하면 합을 반환합니다:

```js
var sumTwoAnd = addTogether(2);
```

`sumTwoAnd(3)`는 `5`를 반환합니다.

어느 인수라도 유효한 숫자가 아니면 undefined를 반환하세요.

# --hints--

`addTogether(2, 3)`는 5를 반환해야 합니다.

```js
assert.deepEqual(addTogether(2, 3), 5);
```

`addTogether(23.4, 30)`는 53.4를 반환해야 합니다.

```js
assert.deepEqual(addTogether(23.4, 30), 53.4);
```

`addTogether("2", 3)`는 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(addTogether('2', 3));
```

`addTogether(5, undefined)`는 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(addTogether(5, undefined));
```

`addTogether("https://www.youtube.com/watch?v=dQw4w9WgXcQ")`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(addTogether('https://www.youtube.com/watch?v=dQw4w9WgXcQ'));
```

`addTogether(5)`는 함수를 반환해야 합니다.

```js
assert.deepEqual(typeof(addTogether(5)), 'function');
```

`addTogether(5)(7)`는 12를 반환해야 합니다.

```js
assert.deepEqual(addTogether(5)(7), 12);
```

`addTogether(2)([3])`는 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(addTogether(2)([3]));
```

`addTogether(2, "3")`는 `undefined`을 반환해야 합니다.

```js
assert.isUndefined(addTogether(2, '3'));
```

# --seed--

## --seed-contents--

```js
function addTogether() {
  return false;
}

addTogether(2,3);
```

# --solutions--

```js
function addTogether() {
  const first = arguments[0];
  if (typeof(first) !== 'number') {
    return undefined;
  }
  if (arguments.length === 1) {
    return function(second) {
      if (typeof(second) !== 'number') {
        return undefined;
      }
      return first + second;
    };
  }
  const second = arguments[1];
  if (typeof(second) !== 'number') {
    return undefined;
  }
  return first + second;
}
```
