---
id: aa7697ea2477d1316795783b
title: 피그 라틴어
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

피그 라틴어는 영어 단어를 변형하는 방법입니다. 규칙은 다음과 같습니다:

\- 단어가 자음으로 시작하면, 첫 자음 또는 자음군을 단어 끝으로 옮기고 `ay`를 추가하세요.

\- 단어가 모음으로 시작하면, 단어 끝에 `way`만 추가하세요.

# --instructions--

제공된 문자열을 피그 라틴어로 변환하세요. 입력 문자열은 모두 소문자인 영어 단어임이 보장됩니다.

# --hints--

`translatePigLatin("california")`는 문자열 `aliforniacay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")`는 문자열 `aragraphspay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")`는 문자열 `oveglay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")`는 문자열 `algorithmway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")`는 문자열 `eightway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

단어 중간에 첫 모음이 나오는 경우도 처리해야 합니다. `translatePigLatin("schwartz")`는 문자열 `artzschway`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

모음이 없는 단어도 처리해야 합니다. `translatePigLatin("rhythm")`는 문자열 `rhythmay`를 반환해야 합니다.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
