---
id: ab306dbdcc907c7ddfc30830
title: 중첩 배열 평탄화
challengeType: 1
forumTopicId: 16079
dashedName: steamroller
---

# --description--

중첩된 배열을 평탄화하세요. 다양한 중첩 수준을 고려해야 합니다.

# --hints--

`steamrollArray([[["a"]], [["b"]]])`는 `["a", "b"]`을 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([[['a']], [['b']]]), ['a', 'b']);
```

`steamrollArray([1, [2], [3, [[4]]]])`는 `[1, 2, 3, 4]`을 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, [2], [3, [[4]]]]), [1, 2, 3, 4]);
```

`steamrollArray([1, [], [3, [[4]]]])`는 `[1, 3, 4]`을 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, [], [3, [[4]]]]), [1, 3, 4]);
```

`steamrollArray([1, {}, [3, [[4]]]])`는 `[1, {}, 3, 4]`을 반환해야 합니다.

```js
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4]);
```

해결책에 `Array.prototype.flat()` 또는 `Array.prototype.flatMap()` 메서드를 사용하지 마세요.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*flat\s*\(/) && !__helpers.removeJSComments(code).match(/\.\s*flatMap\s*\(/));
```

전역 변수는 사용하지 마세요.

```js
steamrollArray([1, {}, [3, [[4]]]])
assert.deepEqual(steamrollArray([1, {}, [3, [[4]]]]), [1, {}, 3, 4])
```

# --seed--

## --seed-contents--

```js
function steamrollArray(arr) {
  return arr;
}

steamrollArray([1, [2], [3, [[4]]]]);
```

# --solutions--

```js
function steamrollArray(arr) {
  if (!Array.isArray(arr)) {
    return [arr];
  }
  var out = [];
  arr.forEach(function(e) {
    steamrollArray(e).forEach(function(v) {
      out.push(v);
    });
  });
  return out;
}
```
