---
id: af7588ade1100bde429baf20
title: 빠진 문자 찾기
challengeType: 1
forumTopicId: 16023
dashedName: missing-letters
---

# --description--

전달된 문자 범위에서 빠진 문자를 찾아 반환하세요.

범위 내 모든 문자가 있으면 `undefined`를 반환하세요.

# --hints--

`fearNotLetter("abce")`는 문자열 `d`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")`는 문자열 `i`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")`는 문자열 `u`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")`는 문자열 `e`를 반환해야 합니다.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")`는 `undefined`를 반환해야 합니다.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js
function fearNotLetter(str) {
  return str;
}

fearNotLetter("abce");
```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    var letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
