---
id: 6565a536ba1f9f25bd30e88b
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

في هذا المشروع، ستنفذ خوارزمية Luhn. هذه الخوارزمية هي صيغة للتحقق من صحة مجموعة متنوعة من أرقام الهوية.

ابدأ بتعريف دالة تسمى `main`، ستكون نقطة الدخول للبرنامج. استخدم الكلمة المفتاحية `pass` لتجنب حدوث خطأ.

# --hints--

يجب أن تحتوي شفرتك على `def main():`.

```js
({
    test: () => {
        assert(runPython(`
        import inspect
        inspect.isfunction(main)
        `))
    }
})
```

يجب أن تحتوي شفرتك على `pass`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const pass = __helpers.python.getDef("\n"+transformedCode, "main");
        assert.include(pass?.function_body, "pass");
    }
})
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
