---
id: 656877c202f44995ed952608
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

استبدل جملة `pass` بمتغير اسمه `card_number` وقيمة `'4111-1111-4555-1142'`.

# --hints--

يجب أن يكون لديك `card_number = '4111-1111-4555-1142'` داخل الدالة `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_number\s*=\s*('|")4111-1111-4555-1142\1/);
    }
})
```

يجب ألا يكون لديك `pass` في شفرتك.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const no_comments = __helpers.python.removeComments(transformedCode);
        assert.notInclude(no_comments, "pass");
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    pass

--fcc-editable-region--
```
