---
id: 656877f71bba2b97acedf9af
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تأتي لغة Python مع أصناف مدمجة تساعدنا في معالجة السلاسل النصية. أحد هذه الأصناف هو `str`. يحتوي على طريقة تسمى `maketrans` تساعدنا في إنشاء جدول ترجمة. يمكن استخدام هذا الجدول لاستبدال الأحرف في سلسلة نصية:

```python
str.maketrans({'t': 'c', 'l': 'b'})
```

الكود أعلاه، عند استدعائه على سلسلة نصية، يستبدل كل الأحرف `t` بـ `c` وكل الأحرف `l` بـ `b`.

أنشئ متغيرًا اسمه `card_translation` وقم بتعيين جدول ترجمة له لاستبدال كل الأحرف `-` و` ` بسلسلة فارغة.

# --hints--

يجب إنشاء متغير `card_translation` داخل الدالة `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +card_translation\s*=/);
    }
})
```

يجب تعيين قيمة `card_translation` إلى `str.maketrans({'-': '', ' ': ''})`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        const allowedMatches = [
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|")-\1\s*:\s*('|")\2\s*,\s*('|") \3\s*:\s*('|")\4\s*\}\s*\)/,
            / +card_translation\s*=\s*str\.maketrans\(\s*\{\s*('|") \1\s*:\s*('|")\2\s*,\s*('|")-\3\s*:\s*('|")\4\s*\}\s*\)/,
        ];
        const anyMatch = allowedMatches.some((match) => match.test(function_body));
        assert.isTrue(anyMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    
--fcc-editable-region--
```
