---
id: 656878585631369a6b2d2191
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

تعريف الترجمة لا يترجم السلسلة النصية بحد ذاته. يجب استدعاء الطريقة `translate` على السلسلة النصية المراد ترجمتها مع جدول الترجمة كمعلمة:

```py
my_string = "tamperlot"
translation_table = str.maketrans({'t': 'c', 'l': 'b'})
translated_string = my_string.translate(translation_table)
```

أنشئ متغيرًا اسمه `translated_card_number` وعيّن له نتيجة استدعاء الطريقة `translate` على `card_number` مع `card_translation` كمعلمة.

# --hints--

يجب إنشاء متغير `translated_card_number` داخل الدالة `main`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=/);
    }
})
```

يجب تعيين قيمة `translated_card_number` إلى `card_number.translate(card_translation)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +translated_card_number\s*=\s*card_number\.translate\(\s*card_translation\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    
--fcc-editable-region--
```
