---
id: 656879c6f35a85a59c06b3a7
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

سبق أن تم الوصول إلى عناصر (حروف) من سلسلة نصية باستخدام عامل الفهرس `[]`. يمكنك أيضًا استخدام عامل الفهرس للوصول إلى نطاق من الحروف في سلسلة نصية باستخدام `string[start:stop:step]`:

```python
my_string = 'camperbot'
my_string[0:6] == 'camper' # True
my_string[0:6:3] == 'cp' # True
```

حيث `start` هو فهرس البداية (شامل)، و`stop` هو فهرس النهاية (غير شامل)، و`step` هو عدد الحروف التي يجب تخطيها. إذا لم يُحدد، يكون `step` افتراضيًا 1.

أنشئ متغيرًا باسم `card_number_reversed` وعيّن له قيمة أول 4 حروف من `card_number`.

# --hints--

يجب أن يحتوي الكود على `card_number_reversed = card_number[0:4]` داخل الدالة `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /card_number_reversed\s*=\s*card_number\s*\[\s*0\s*:\s*4\s*\]/);
    }
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
