---
id: 65687b48f2201ab32e06c37c
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

حاليًا، يُصدر السكربت الخاص بك خطأً من نوع `TypeError` لأنك تحاول إضافة سلسلة نصية إلى عدد صحيح. يمكنك إصلاح ذلك بتحويل المتغير `digit` إلى عدد صحيح قبل إضافته إلى `sum_of_odd_digits`، باستخدام الدالة المدمجة `int`:

```python
my_string = '123'
my_int = int(my_string)
```

حوّل المتغير `digit` إلى عدد صحيح قبل إضافته إلى `sum_of_odd_digits`. ثم انقل استدعاء `print` إلى نهاية دالة `verify_card_number` لطباعة قيمة `sum_of_odd_digits`.

# --hints--

يجب أن يحتوي الحلقة `for` على `sum_of_odd_digits += int(digit)`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*\+=\s*int\(\s*digit\s*\)/);
    }
})
```

يجب أن يكون لديك `print(sum_of_odd_digits)` في نهاية دالة `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /print\(\s*sum_of_odd_digits\s*\)/);
    }
})
```

يجب ألا يكون لديك `print(digit)` بعد الآن.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        const no_comments = __helpers.python.removeComments(function_body);

        assert.notMatch(no_comments, /print\(\s*digit\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

--fcc-editable-region--
    for digit in odd_digits:
        print(digit)
        sum_of_odd_digits += digit
        
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
