---
id: 656879a66338b2a461d5d307
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

<!-- TODO: Find better places to split explanation up. -->

خوارزمية Luhn هي كما يلي:

1. من اليمين إلى اليسار، ضاعف قيمة كل رقم ثاني؛ إذا كان الناتج أكبر من 9، اجمع أرقام الناتج.
2. خذ مجموع كل الأرقام.
3. إذا كان مجموع كل الأرقام من مضاعفات 10، فالرقم صحيح؛ وإلا فهو غير صحيح.

افترض مثالًا لرقم حساب "7992739871" سيتم إضافة رقم تحقق له، ليصبح بالشكل 7992739871x:

```markdown
Account number      7   9  9  2  7  3  9   8  7  1  x
Double every other  7  18  9  4  7  6  9  16  7  2  x
Sum 2-char digits   7   9  9  4  7  6  9   7  7  2  x
```

استبدل عبارة `pass` بمتغير `sum_of_odd_digits` وقيمة `0`.

# --hints--

يجب أن يكون لديك `sum_of_odd_digits = 0` داخل دالة `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.match(function_body, /sum_of_odd_digits\s*=\s*0/);
    }
})
```

يجب ألا يكون لديك `pass` داخل دالة `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        const { function_body } = verify_card_number;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def verify_card_number(card_number):
    pass

--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

    verify_card_number(translated_card_number)

main()
```
