---
id: 65687db6a9fcf3c548a6c1b7
title: الخطوة 28
challengeType: 20
dashedName: step-28
---

# --description--

جزء من الخوارزمية هو مضاعفة كل رقم ثاني، بدءًا من اليمين. إذا كانت نتيجة مضاعفة الرقم أكبر من أو تساوي `10`، اجمع الرقمين معًا. على سبيل المثال، إذا كان الرقم `6`، ضاعفه لتحصل على `12`. اجمع `1` و `2` لتحصل على `3`. يمكنك فعل ذلك باستخدام _القسمة الصحيحة_ للحصول على الرقم الأول ومعامل باقي القسمة (`%`) للحصول على الرقم الثاني:

```python
my_number = 12
first_digit = my_number // 10
second_digit = my_number % 10
```

القسمة الصحيحة تُنتج حاصل القسمة، مقربًا إلى أقرب عدد صحيح أقل أو يساوي الناتج.

داخل جملة `if`، عيّن لـ `number` ناتج `number // 10` (القسمة الصحيحة) زائد معامل باقي القسمة بين `number` و `10`.

# --hints--

يجب أن تحتوي جملة `if` على `number = number // 10 + number % 10`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getBlock("\n" + transformedCode, "if number >= 10");
        const { block_body } = verify_card_number;
        const allowedMatches = [
            /number *= *\(?\s*number *\/\/ *10\s*\)? *\+ *\(?\s*number *% *10/,
            /number *= *\(?\s*number *% *10\s*\)? *\+ *\(?\s*number *\/\/ *10/
        ];

        const someMatch = allowedMatches.some(allowedMatch => allowedMatch.test(block_body));
        assert.isTrue(someMatch);
    }
})
```

# --seed--

## --seed-contents--

```py
def verify_card_number(card_number):
    sum_of_odd_digits = 0
    card_number_reversed = card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

--fcc-editable-region--
    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            print(number)
            
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    verify_card_number(translated_card_number)

main()
```
