---
id: 6568789edf2ed39c81983cc4
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

اطبع رقم البطاقة المترجم في وحدة التحكم.

# --hints--

يجب أن تطبع `translated_card_number` في وحدة التحكم.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*translated_card_number\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)
    
--fcc-editable-region--
```
