---
id: 6569b831a07d7154c793301b
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

عرّف دالة `verify_card_number` بمعلمة `card_number`، واستخدم الكلمة المفتاحية `pass` لجعل الدالة لا تفعل شيئًا.

# --hints--

يجب أن يكون لديك دالة تسمى `verify_card_number` بمعلمة تسمى `card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_parameters } = verify_card_number;
        assert.equal(function_parameters.trim(), "card_number");
    }
})
```

يجب استخدام الكلمة المفتاحية `pass` في جسم دالة `verify_card_number`.

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const verify_card_number = __helpers.python.getDef("\n" + transformedCode, "verify_card_number");
        assert.exists(verify_card_number);
        const { function_body } = verify_card_number;
        assert.match(function_body, / +pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

def main():
    card_number = '4111-1111-4555-1142'
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    print(translated_card_number)

main()
```
