---
id: 68f890b73bbfc9efa7379e7b
title: Налагодження помилок перетворення типів у проблемному застосунку
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

Ви щойно приєдналися до місцевої команди веброзробки, і ваше перше завдання — виправити проблемний код, залишений попередніми розробниками.

Вони намагалися виконати прості арифметичні дії, але щось пішло не так, і результати не мають сенсу. Ваше завдання — переглянути, налагодити та виправити код, щоб він працював як очікується і був легшим для читання.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути змінна з іменем `firstResult`, яка правильно додає числа `5` і `10`, щоб отримати значення `15`.
2. У вас має бути змінна з іменем `secondResult`, яка правильно віднімає `5` від `8`, щоб отримати значення `3`.
3. У вас має бути змінна з іменем `thirdResult`, яка отримує значення `6` шляхом додавання двох чисел. Замініть булеве значення, яке зараз використовується у виразі, на число.
4. У вас має бути змінна з іменем `fourthResult`, яка отримує значення `8` шляхом додавання двох чисел. Замініть булеве значення, яке зараз використовується у виразі, на число.
5. У вас має бути змінна з іменем `fifthResult`, яка правильно множить два числа, щоб отримати добуток `20`. Замініть рядок, який зараз використовується у виразі, на число.
6. У вас має бути змінна з іменем `sixthResult`, яка правильно додає два числа, щоб отримати значення `22`. Замініть `null`, яке зараз використовується у виразі, на число.

# --hints--

У вас має бути змінна з іменем `firstResult`.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` має мати числове значення `15`.

```js
assert.strictEqual(firstResult, 15);
```

У вас має бути змінна з іменем `secondResult`.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` має мати числове значення `3`.

```js
assert.strictEqual(secondResult, 3);
```

У вас має бути змінна з іменем `thirdResult`.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` має мати числове значення `6`.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` не має містити булеве значення (`true` або `false`).

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

У вас має бути змінна з іменем `fourthResult`.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` має мати числове значення `8`.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` не має містити булеве значення (`true` або `false`).

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

У вас має бути змінна з іменем `fifthResult`.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` має мати числове значення `20`.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` не має містити множення рядка (наприклад, "10" * 2).

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

У вас має бути змінна з іменем `sixthResult`.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` має мати числове значення `22`.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` не має містити `null` у своєму виразі.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
