---
id: 669aff9f5488f1bea056416d
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with semantic HTML by building a blog page dedicated to Mr. Whiskers the cat.

Para iniciar o projeto, adicione o `<!DOCTYPE html>` e um elemento `html` com um atributo `lang` de `en`.

Remember that you learned how to build a basic HTML boilerplate like this in the previous module.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

Você deve ter a declaração `<!DOCTYPE html>`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Você deve ter uma tag de abertura para `html` com o idioma definido para inglês.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Você deve acrescentar uma tag de fechamento para `html`.

```js
assert.match(code, /<\/html>/i);
```

O `DOCTYPE` deve vir antes do elemento `html`.

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
