---
id: 669fdfc9a5e3da42d2376609
title: Passo 13
challengeType: 0
dashedName: step-13
---

# --description--

Agora que você adicionou a seção sobre, tente clicar no link de About para ver a página descer até essa seção.

A próxima seção na página do blog será uma lista de postagens falando sobre o sr. Whiskers.

Adicione outro elemento `section` com um `id` definido como `"posts"`.

Dentro do elemento `section`, adicione um elemento `h2` com o texto `Posts`.

# --hints--

O código deve ter uma tag de abertura para `section`.

```js
assert.match(code, /<section.*>/i);
```

O código deve ter uma tag de fechamento para `section`.

```js
assert.match(code, /<\/section>/i);
```

O elemento `section` deve ter um `id` definido como `"posts"`.

```js
assert.match(code, /<section.*id="posts".*>/i);
```

O elemento `section` deve estar dentro do elemento `main`.

```js
const sectionElement = document.getElementById('posts');
assert.strictEqual(sectionElement.parentElement.tagName, 'MAIN');
```

O elemento `h2` deve estar dentro do elemento `section`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element.parentElement.tagName, 'SECTION');
```

O elemento `h2` deve ter o texto `Posts`.

```js
const h2Element = document.querySelector('#posts h2');
assert.strictEqual(h2Element.innerText, 'Posts');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
      --fcc-editable-region--

      --fcc-editable-region--
    </main>
  </body>
</html>
```
