---
id: 669fd6fd12918e3de87854d4
title: Passo 8
challengeType: 0
dashedName: step-8
---

# --description--

Para o seu blog, deve haver uma maneira para os usuários navegarem para diferentes seções na página.

The `nav` element is used to provide navigation links to other sections in the document or other sections in the website. Muitas vezes, você verá o elemento `nav` usado para menus ou tabela de conteúdos.

Aqui está um exemplo do uso do elemento `nav`:

```html
<nav>
  <ul>
    <li><a href="#">Home</a></li>
    <li><a href="#about">About</a></li>
    <li><a href="#contact">Contact</a></li>
  </ul>
</nav>
```

Abaixo do elemento `figure`, adicione um elemento `nav` com um elemento `ul` aninhado dentro dele.

Dentro do elemento `ul`, adicione três elementos `li`.

# --hints--

O elemento `nav` deve ter uma tag de abertura.

```js
assert.match(code, /<nav\>/);
```

O elemento `nav` deve ter uma tag de fechamento.

```js
assert.match(code, /<\/nav\>/);
```

O elemento `nav` deve estar dentro do elemento `header`.

```js
const navElement = document.querySelector('nav');
assert.strictEqual(navElement?.parentElement.tagName, 'HEADER');
```

O elemento `ul` deve ter uma tag de abertura.

```js
assert.match(code, /<ul\>/);
```

O elemento `ul` deve ter uma tag de fechamento.

```js
assert.match(code, /<\/ul\>/);
```

O elemento `ul` deve estar dentro do elemento `nav`.

```js
const ulElement = document.querySelector('ul');
assert.strictEqual(ulElement?.parentElement.tagName, 'NAV');
```

O elemento `ul` deve conter três elementos `li`.

```js
const ulElement = document.querySelector("ul");

const liElements = ulElement?.querySelectorAll("li");
assert.strictEqual(liElements?.length, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat in the garden"
        />
        <figcaption>Mr. Whiskers in the Garden</figcaption>
      </figure>
--fcc-editable-region--

--fcc-editable-region--
    </header>
  </body>
</html>
```
