---
id: 669fc938d38e6e38ace9251e
title: Passo 3
challengeType: 0
dashedName: step-3
---

# --description--

Dentro do elemento `head`, aninhe um elemento `meta` com o atributo `charset` definido como `"UTF-8"`.

Abaixo desse elemento `meta`, adicione um elemento `title`.

O texto do elemento `title` deve ser `Mr. Whiskers' Blog`.

# --hints--

Você deve ter um elemento `meta`.

```js
assert.isNotNull(document.querySelector("meta"));
```

O elemento `meta` deve ser um elemento nulo, sem uma tag final `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

A tag `meta` deve ter um atributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

O atributo `charset` deve ter um valor de `"UTF-8"`.

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

O elemento `meta` deve estar aninhado dentro do elemento `head`.

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

Você deve acrescentar uma tag de abertura`title`.

```js
assert.match(code, /<title>/i);
```

Você deve acrescentar uma tag de fechamento para `title`.

```js
assert.match(code, /<\/title>/i);
```

O elemento `title` deve estar dentro do elemento `head`.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

O elemento `title` deve ter o texto `Mr. Whiskers' Blog`. Não se esqueça de verificar a ortografia.

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Mr. Whiskers' Blog");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
