---
id: 66a334a1a7cca6354999f9bf
title: Passo 14
challengeType: 0
dashedName: step-14
---

# --description--

Para a primeira postagem no blog, você usará um elemento <dfn>article</dfn>.

O elemento `article` representa conteúdo independente em uma página da web.

```html
<article>
  <h1>Example heading</h1>
  <p>Example article text</p>
</article>
```

Abaixo do elemento `h2`, adicione um elemento `article`.

Dentro do elemento `article`, adicione um elemento `h3` com o texto `Mr. Whiskers' First Day Home`.

A razão pela qual um `h3` é usado aqui está na importância de se manter uma hierarquia estrutural adequada para os elementos de cabeçalho. Como o subtítulo da postagem é um elemento `h2`, o próximo nível na hierarquia seria um `h3`.

# --hints--

O código deve ter uma tag de abertura `article`.

```js
assert.match(code, /<article>/i);
```

O código deve ter uma tag de fechamento `article`.

```js
assert.match(code, /<\/article>/i);
```

O elemento `article` deve estar dentro do elemento `section` com o `id` definido como `posts`.

```js
const articleElement = document.querySelector('#posts article');
assert.strictEqual(articleElement?.parentElement.tagName, 'SECTION');
```

O código deve ter uma tag de abertura `h3`.

```js
assert.match(code, /<h3>/i);
```

O código deve ter uma tag de fechamento `h3`.

```js
assert.match(code, /<\/h3>/i);
```

O elemento `h3` deve ter o texto de `Mr. Whiskers' First Day Home`.

```js
const h3Element = document.querySelector('h3');
assert.strictEqual(h3Element?.innerText.trim(), "Mr. Whiskers' First Day Home");
```

O elemento `h3` deve estar dentro do elemento `article`.

```js
const h3Element = document.querySelector('#posts article h3');
assert.strictEqual(h3Element?.parentElement.tagName, 'ARTICLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      <section id="about">
        <h2>About</h2>
        <p>
          Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.
        </p>
        <p>
          His playful nature and boundless energy keeps me on my toes. I love him so much.
        </p>
      </section>
--fcc-editable-region--
      <section id="posts">
        <h2>Posts</h2>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
