---
id: 66a49f685961e997e337cab1
title: Passo 6
challengeType: 0
dashedName: step-6
---

# --description--

O `header` será responsável por exibir o título principal, a imagem e a barra de navegação da página para o blog.

Dentro do elemento `header`, adicione um `h1` com o texto `Welcome to Mr. Whiskers' Blog Page!`.

# --hints--

Você deve acrescentar uma tag de abertura para `h1`.

```js
assert.match(code, /<h1>/i);
```

Você deve acrescentar uma tag de fechamento para `h1`.

```js
assert.match(code, /<\/h1>/i);
```

O elemento `h1` deve ter o texto `Welcome to Mr. Whiskers' Blog Page!`. Verifique novamente se há erros de ortografia.

```js
const h1Text = document.querySelector("h1")?.innerText
assert.strictEqual(h1Text, "Welcome to Mr. Whiskers' Blog Page!");
```

O elemento `h1` deve estar dentro do elemento `header`.

```js
const h1Element = document.querySelector("h1");
assert.strictEqual(h1Element?.parentElement.tagName, 'HEADER');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    <header>

    </header>
--fcc-editable-region--
  </body>
</html>
```
