---
id: 66a7e72adf226c02626715a3
title: Passo 12
challengeType: 0
dashedName: step-12
---

# --description--

Abaixo do elemento `h2`, adicione um elemento de parágrafo com o texto de `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.`.

Abaixo do elemento de parágrafo, adicione outro elemento de parágrafo com o texto `His playful nature and boundless energy keeps me on my toes. I love him so much.`

# --hints--

O primeiro parágrafo deve ter o texto de `Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers.` Não se esqueça de verificar a ortografia.

```js
const pElement = document.querySelector('body p:first-of-type');
const text = "Hi there! I'm Jane Doe, a passionate writer who finds endless inspiration in the antics of my beloved cat, Mr. Whiskers."
assert.strictEqual(pElement?.innerText, text);
```

O primeiro parágrafo deve estar aninhado dentro do elemento `section`.

```js
const pElement = document.querySelector('#about p:first-of-type');
assert.strictEqual(pElement?.parentElement.tagName, 'SECTION');
```

Você deve ter um segundo parágrafo dentro do elemento `section`.

```js
const pElement = document.querySelectorAll('#about p');
assert.strictEqual(pElement?.length, 2);
```

O segundo parágrafo deve ter o texto `His playful nature and boundless energy keeps me on my toes. I love him so much.` Não se esqueça de verificar a ortografia.

```js
const pElement = document.querySelector('#about p:last-of-type');
const text = "His playful nature and boundless energy keeps me on my toes. I love him so much."
assert.strictEqual(pElement?.innerText, text);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Mr. Whiskers' Blog</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <header>
      <h1>Welcome to Mr. Whiskers' Blog Page!</h1>
      <figure>
        <img
          src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg"
          alt="a cat peacefully sleeping"
        />
        <figcaption>Mr. Whiskers Sleeping</figcaption>
      </figure>
      <nav>
        <ul>
          <li><a href="#about">About</a></li>
          <li><a href="#posts">Posts</a></li>
          <li><a href="#contact">Contact</a></li>
        </ul>
      </nav>
    </header>
    <main>
      --fcc-editable-region--
      <section id="about">
        <h2>About</h2>

      </section>
      --fcc-editable-region--
    </main>
  </body>
</html>
```
