---
id: 68599af28ead38c22590a89d
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Потрібен спосіб зв’язати кнопки з вмістом, який вони представляють. Це робиться за допомогою атрибута `aria-controls` на кожній кнопці. Пізніше ви додасте вміст і відповідні атрибути `id` до елементів вмісту.

Почніть із додавання атрибута `aria-controls` до кожного з елементів `button`. Значення атрибута `aria-controls` першої кнопки має бути встановлене на `"panel-earth"`. Атрибут `aria-controls` другої кнопки має бути встановлений на `"panel-saturn"`. Атрибут `aria-controls` третьої кнопки має бути встановлений на `"panel-mars"`.

# --hints--

Кожен з ваших елементів `button` має мати атрибут `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

Перший елемент `button` має мати атрибут `aria-controls`, встановлений на `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

Другий елемент `button` має мати атрибут `aria-controls`, встановлений на `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

Третій елемент `button` має мати атрибут `aria-controls`, встановлений на `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
