---
id: 68599633407374c0d574587b
title: Крок 3
challengeType: 0
dashedName: step-3
---

# --description--

Наступний крок — додати кнопки, які відповідатимуть за перемикання між різними планетами та їхніми фактами.

Всередині елемента `tablist` додайте три елементи `button`, кожен з атрибутом `role`, встановленим у `"tab"`. Перший елемент `button` повинен містити текст `🌍 Earth`. Другий елемент `button` повинен містити текст `🪐 Saturn`. Третій елемент `button` повинен містити текст `🔴 Mars`.  

# --hints--

У вас має бути три елементи `button` всередині елемента `tablist`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

Кожен з ваших елементів `button` повинен мати атрибут `role`, встановлений у `"tab"`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

Перший елемент `button` повинен містити текст `🌍 Earth`.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

Другий елемент `button` повинен містити текст `🪐 Saturn`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

Третій елемент `button` повинен містити текст `🔴 Mars`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./style.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
