---
id: 685976a92ffab8b2dc744b72
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Всередині вашого елемента `.tabs` додайте елемент `div` з атрибутом `role`, встановленим у `"tablist"`, та атрибутом `aria-labelledby`, встановленим у `"tabs-title"`.  

Атрибут `aria-labelledby` потрібен тут, щоб пов’язати tablist із видимою міткою, допомагаючи допоміжним технологіям, таким як читач екрана, оголосити користувачам призначення tablist.

# --hints--

Усередині вашого елемента `div` має бути елемент `.tabs`.

```js
assert.exists(document.querySelector(".tabs div"));
```

Ваш елемент `div` має мати атрибут `role`, встановлений у `"tablist"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("role"), "tablist");
```

Ваш елемент `div` має мати атрибут `aria-labelledby`, встановлений у `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs div").getAttribute("aria-labelledby"), "tabs-title");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>

      --fcc-editable-region--

      --fcc-editable-region--
    </div>
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
