---
id: 685db0bd1280154ae85bcf9f
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

У цьому практичному занятті ви створите доступний список вкладок планет за допомогою HTML, CSS і JavaScript. Шаблонний код HTML і CSS вже надано.

На цих перших кроках ви створите решту HTML для списку вкладок.

Почніть зі створення елемента `div` з множиною класів `"tabs"`. Всередині цього елемента `div` додайте елемент `h2` з атрибутом `id`, встановленим у `"tabs-title"`, і текстовим вмістом `Planets`.

# --hints--

У вас має бути елемент `div`.

```js
assert.exists(document.querySelector("div"));
```

Ваш елемент `div` має мати множину класів `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

Ваш елемент `div` має містити всередині елемент `h2`.

```js
assert.exists(document.querySelector(".tabs h2"));
```

Ваш елемент `h2` має мати атрибут `id`, встановлений у `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

Ваш елемент `h2` має містити текстовий вміст `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
