---
id: 66068fb0bfddba2b7977eb60
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você aprenderá sobre classes e objetos construindo um resolvedor de quebra-cabeça sudoku.

Em Python, uma classe é um modelo para criar objetos. Objetos criados a partir de uma classe são instâncias dessa classe. Você pode criar uma classe usando esta sintaxe:

```py
class ClassName:
    pass
```

Onde `class` é a palavra-chave necessária para definir a classe e `ClassName` é o nome da classe, escrito por convenção em *PascalCase*.

Comece criando uma classe `Board`.

# --hints--

Você deve criar uma classe chamada `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
