---
id: 66069f86f58f85371d47123e
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Um atributo é uma variável associada a um objeto, que é usada para armazenar dados como variáveis regulares.

Dentro do método `__init__`, atribua o parâmetro `board` (que é passado ao criar uma instância da classe `Board`) a um atributo de instância `board` usando `self.board`.

`self.board` refere-se ao atributo `board` da instância da classe. É uma variável que pertence ao objeto criado a partir da classe `Board`.

# --hints--

Você deve deletar `pass` e atribuir o parâmetro `board` a `self.board` dentro do método `__init__`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").find_body().is_equivalent("self.board = board")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def __init__(self, board):
        pass

puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
--fcc-editable-region--
```
