---
id: 6606927d010be4300a4e5330
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

The instantiation creates an empty object.Mas classes podem ter métodos, que são como funções locais para cada instância.Dentro de uma classe, os métodos são declarados da seguinte forma:

```py
class ClassName:
    def method_name():
        pass
```

Dentro da classe `Board`, substitua `pass` por um método vazio `spam`.

# --hints--

Você deve definir um método chamado `spam` dentro da classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
    
gameboard = Board()
--fcc-editable-region--
```
