---
id: 660699119472f332798860ad
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Agora, substitua `pass` por uma chamada `print` e passe a ela a string `'Spam!'`.

# --hints--

Você deve deletar `pass` e imprimir a string `'Spam!'` dentro do método `spam`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").find_body().is_equivalent("print('Spam!')")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        pass
    
gameboard = Board()
--fcc-editable-region--
```
