---
id: 66069167b3307b2f4067b22b
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Uma nova instância de uma classe é criada usando a notação de função, que envolve adicionar um par de parênteses ao nome da classe.

Fora da definição da classe, crie uma instância da classe `Board` e atribua-a a uma variável chamada `gameboard`.

# --hints--

Você deve declarar uma variável `gameboard` fora da classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("gameboard")`)) })
```

Sua variável `gameboard` deve ter o valor de `Board()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("gameboard").is_equivalent("gameboard = Board()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    pass
--fcc-editable-region--
```
