---
id: 6606c14182435d4bab0de2ee
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

Em seguida, você vai criar um método que verifica se um número pode ser inserido em uma coluna especificada do tabuleiro de sudoku verificando se o número não está presente nessa coluna.

Dentro da classe `Board`, crie um método chamado `valid_in_col` e dê a ele três parâmetros: `self`, `col` e `num`. Onde `col` e `num` são, respectivamente, o índice da coluna e o número a ser verificado.

# --hints--

Você deve criar um novo método chamado `valid_in_col` dentro da classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_col")`)) })
```

Seu método `valid_in_col` deve ter três parâmetros: `self`, `col` e `num`, nesta ordem.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_col").has_args("self, col, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
