---
id: 6606933d6813a8308c962dd1
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Para ser um método de instância, um método requer um parâmetro especial, chamado `self` por convenção.Este parâmetro é uma referência à instância da classe e deve sempre ser o primeiro parâmetro.

Adicione um parâmetro `self` ao seu método `spam`.

# --hints--

Seu método `spam` deve ter um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("spam").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam():
        pass
    
gameboard = Board()
--fcc-editable-region--
```
