---
id: 6606d8c323d6205890fbbd54
title: Passo 46
challengeType: 20
dashedName: step-46
---

# --description--

Dentro da classe `Board`, crie outro método `is_valid` e dê a ele três parâmetros: `self`, `empty` e `num`.Onde `empty` é uma tupla que representa os índices da linha e da coluna de uma célula vazia e `num` é o número a ser verificado.

Este método verificará se um número dado é uma escolha válida para uma célula vazia no tabuleiro de sudoku, validando sua compatibilidade com a linha, coluna e o quadrado 3x3 da célula vazia especificada.

# --hints--

Você deve criar um novo método chamado `is_valid` dentro da classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("is_valid")`)) })
```

Seu método `is_valid` deve ter três parâmetros: `self`, `empty` e `num`, nesta ordem.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("is_valid").has_args("self, empty, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None

    def valid_in_row(self, row, num):
        return num not in self.board[row]

    def valid_in_col(self, col, num):
        return all(self.board[row][col] != num for row in range(9))

    def valid_in_square(self, row, col, num):
        row_start = (row // 3) * 3
        col_start = (col // 3) * 3
        for row_no in range(row_start, row_start + 3):
            for col_no in range(col_start, col_start + 3):
                if self.board[row_no][col_no] == num:
                    return False
        return True
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
