---
id: 660699aabc59c532f2d556e5
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Para chamar um método de instância, você precisa usar a notação de ponto:

```py
instance_name.method_name()
```

Onde `instance_name` é a instância ou objeto e `method_name` é o método que você quer chamar.

Chame o método `spam` do objeto `gameboard`.

# --hints--

Você deve chamar o método `spam` do objeto `gameboard` com `gameboard.spam()`.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("gameboard.spam()")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
