---
id: 6606a4641ec48b3a9fe8c2fc
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

A função embutida `enumerate` recebe um iterável como argumento e retorna um objeto enumerate que você pode iterar.Ele fornece a contagem (que por padrão começa em zero) e o valor do iterável.

```py
iterable = ['a', 'b', 'c']
for i, j in enumerate(iterable):
    print(i, j)
```

O loop do exemplo acima exibiria as tuplas `0, a`, `1, b` e `2, c`.

Dentro do método `find_empty_cell`, substitua `pass` por um loop `for` que usa a função `enumerate()` para iterar sobre cada linha no tabuleiro de sudoku.Use `row` como o índice da linha atual e `contents` para os elementos da linha atual.


# --hints--

Você deve substituir `pass` por um loop `for` que itere sobre `enumerate(self.board)`.Use `row` e `contents` como as variáveis do loop.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("find_empty_cell").find_body().is_equivalent("for row, contents in enumerate(self.board):\\n  pass")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board
--fcc-editable-region--
    def find_empty_cell(self):
        pass
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
