---
id: 66069b992c1c5e3451f3deb0
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

The instantiation creates an empty object.O método `__init__` é um método especial que permite instanciar um objeto em um estado personalizado.Quando uma classe implementa um método `__init__`, `__init__` é chamado automaticamente na instanciação.

Dentro da sua classe `Board`, exclua o método `spam` e substitua-o por um método `__init__` que inclua um parâmetro `self`.

# --hints--

Você não deve ter um método `spam` na sua classe `Board`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("Board").has_function("spam")`)) })
```

Você deve definir um método `__init__` na sua classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("__init__")`)) })
```

Seu método `__init__` deve ter um parâmetro `self`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("__init__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Board:
    def spam(self):
        print('Spam!')
    
gameboard = Board()
--fcc-editable-region--
```
