---
id: 6606bc4e5535c0484990ccd5
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Em seguida, você vai trabalhar em um método que verifica se um determinado número pode ser inserido em uma linha especificada do tabuleiro de sudoku.

Dentro da classe `Board`, crie um método chamado `valid_in_row` e dê a ele três parâmetros: `self`, `row` e `num`. Onde `self` representa a instância da classe e `row` e `num` são o índice da linha e o número a ser verificado, respectivamente.


# --hints--

Você deve criar um novo método chamado `valid_in_row` dentro da classe `Board`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").has_function("valid_in_row")`)) })
```

Seu método `valid_in_row` deve ter três parâmetros: `self`, `row` e `num`, nesta ordem.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Board").find_function("valid_in_row").has_args("self, row, num")`)) })
```

# --seed--

## --seed-contents--

```py
class Board:
    def __init__(self, board):
        self.board = board

    def find_empty_cell(self):
        for row, contents in enumerate(self.board):
            try:
                col = contents.index(0)
                return row, col
            except ValueError:
                pass
        return None
--fcc-editable-region--
    
--fcc-editable-region--
puzzle = [
  [0, 0, 2, 0, 0, 8, 0, 0, 0],
  [0, 0, 0, 0, 0, 3, 7, 6, 2],
  [4, 3, 0, 0, 0, 0, 8, 0, 0],
  [0, 5, 0, 0, 3, 0, 0, 9, 0],
  [0, 4, 0, 0, 0, 0, 0, 2, 6],
  [0, 0, 0, 4, 6, 7, 0, 0, 0],
  [0, 8, 6, 7, 0, 4, 0, 0, 0],
  [0, 0, 0, 5, 1, 9, 0, 0, 8],
  [1, 7, 0, 0, 0, 6, 0, 0, 5]
]

gameboard = Board(puzzle)
```
