---
id: 579e2a2c335b9d72dd32e05c
title: Tekeleza algoriti ya kata na ungana
challengeType: 26
dashedName: implement-the-slice-and-splice-algorithm
---

# --description--

Katika maabara hii, utahitaji kuunda algoriti ya kuunganisha safu mbili za data.

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unda kitendakazi cha `frankenSplice` kinachokubali safu mbili za data na kielezo.
1. Nakili kila kipengele cha safu ya kwanza ndani ya safu ya pili, kwa mpangilio, kuanzia kwenye kielezo kilichotolewa, na rudisha safu inayotokana.
1. Safu za data za ingizo ziendelee kuwa zile zile baada ya kitendakazi kuendeshwa.

# --hints--

`frankenSplice([1, 2, 3], [4, 5], 1)` inapaswa kurudisha `[4, 1, 2, 3, 5]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3], [4, 5], 1), [4, 1, 2, 3, 5]);
```

`frankenSplice([1, 2], ["a", "b"], 1)` inapaswa kurudisha `["a", 1, 2, "b"]`.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
assert.deepEqual(frankenSplice(testArr1, testArr2, 1), ['a', 1, 2, 'b']);
```

`frankenSplice(["claw", "tentacle"], ["head", "shoulders", "knees", "toes"], 2)` inapaswa kurudisha `["head", "shoulders", "claw", "tentacle", "knees", "toes"]`.

```js
assert.deepEqual(
  frankenSplice(
    ['claw', 'tentacle'],
    ['head', 'shoulders', 'knees', 'toes'],
    2
  ),
  ['head', 'shoulders', 'claw', 'tentacle', 'knees', 'toes']
);
```

Vipengele vyote kutoka safu ya kwanza vinapaswa kuongezwa kwenye safu ya pili kwa mpangilio wao wa awali. `frankenSplice([1, 2, 3, 4], [], 0)` inapaswa kurudisha `[1, 2, 3, 4]`.

```js
assert.deepEqual(frankenSplice([1, 2, 3, 4], [], 0), [1, 2, 3, 4]);
```

Safu ya kwanza inapaswa kubaki ile ile baada ya kitendakazi kuendeshwa.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr1, [1, 2]);
```

Safu ya pili inapaswa kubaki ile ile baada ya kitendakazi kuendeshwa.

```js
let testArr1 = [1, 2];
let testArr2 = ['a', 'b'];
frankenSplice(testArr1, testArr2, 1);
assert.deepEqual(testArr2, ['a', 'b']);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function frankenSplice(arr1, arr2, n) {
  let result = arr2.slice();
  for (let i = 0; i < arr1.length; i++) {
    result.splice(n + i, 0, arr1[i]);
  }
  return result;
}
```
