---
id: 66b102ee0bdbad34a9f85ba0
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop continuerai a imparare sulle stringhe costruendo un Teacher Chatbot.

Per cominciare, aggiungi un'istruzione `console` con il messaggio `"Hi there!"`.

# --hints--

Dovresti avere un'istruzione `console`.

```js
assert.match(code, /console\.log(.*)/);
```

La tua istruzione console dovrebbe mostrare il messaggio `"Hi there!"`.

```js
assert.match(code, /console\.log\((['"])(Hi\s+there!)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
