---
id: 66b59829dba144ff1351220f
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ora è il momento di creare un saluto usando la variabile `botName`.

Nelle lezioni precedenti, hai imparato a concatenare stringhe usando i template literal in questo modo:

```js
const name = "John";

// "Hello, John!"
`Hello, ${name}!`;
```

Inizia creando una variabile chiamata `greeting`.

Poi, usando la sintassi dei template literal, assegna una stringa che dice `My name is`, seguita dalla variabile `botName`, e termina con un punto (`.`).

Infine, stampa la variabile `greeting` nella console.

# --hints--

Dovresti avere una variabile chiamata `greeting`.

```js
assert.isNotNull(greeting);
```

Dovresti assegnare una stringa alla tua variabile `greeting`.

```js
assert.isString(greeting);
```

Dovresti usare i template literal per concatenare la stringa `My name is` con la variabile `botName` seguita da un punto (`.`).

```js
assert.equal(greeting, "My name is teacherBot.");
```

Dovresti avere un’istruzione `console` che stampa la variabile `greeting` nella console.

```js
assert.match(code, /console\.log\(.*greeting.*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
