---
id: 66b59b12a745e10011158f7b
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Il passaggio successivo è creare alcune variabili in più che saranno usate nei messaggi futuri del bot.

Crea una variabile chiamata `subject` e assegnale il valore stringa `"JavaScript"`.

Poi crea una variabile chiamata `topic` e assegnale il valore stringa `"strings"`.

# --hints--

Dovresti avere una variabile chiamata `subject`.

```js
assert.isNotNull(subject);
```

Dovresti assegnare una stringa alla tua variabile `subject`.

```js
assert.isString(subject);
```

Dovresti assegnare la stringa `"JavaScript"` alla tua variabile `subject`.

```js
assert.equal(subject, "JavaScript");
```

Dovresti avere una variabile chiamata `topic`.

```js
assert.isNotNull(topic);
```

Dovresti assegnare una stringa alla tua variabile `topic`.

```js
assert.isString(topic);
```

Dovresti assegnare la stringa `"strings"` alla tua variabile `topic`.

```js
assert.equal(topic, "strings");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

--fcc-editable-region--

--fcc-editable-region--
```
