---
id: 66b59be6ab830800c4df9146
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora è il momento di usare le variabili che hai creato nel passaggio precedente.

Nelle lezioni precedenti, hai imparato a lavorare con i template literal in questo modo:

```js
const name = "John";
const age = 30;

// My name is John and I am 30 years old.
`My name is ${name} and I am ${age} years old.`;
```

Inizia creando una variabile `sentence`.

Usando la sintassi dei template literal, assegna la stringa `Today, you will learn about [topic variable goes here] in [subject variable goes here].` alla variabile `sentence`.

Sostituirai i segnaposto `[topic variable goes here]` e `[subject variable goes here]` con le variabili `topic` e `subject` e la sintassi `${}`.

Infine, stampa la variabile `sentence` nella console.

# --hints--

Dovresti avere una variabile `sentence`.

```js
assert.isNotNull(sentence);
```

La tua variabile `sentence` dovrebbe essere una stringa.

```js
assert.isString(sentence);
```

Dovresti usare la sintassi dei template literal per assegnare la stringa `Today, you will learn about [topic variable goes here] in [subject variable goes here].` alla variabile `sentence`. Consulta l’esempio se hai bisogno di aiuto.

```js
assert.match(code, /Today,\s*you\s+will\s+learn\s+about\s+\$\{topic\}\s+in\s+\$\{subject\}\./);
```

Dovresti stampare la variabile `sentence` nella console.

```js
assert.match(code,/console\.log\(\s*sentence\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

--fcc-editable-region--

--fcc-editable-region--
```
