---
id: 66b5928b86e907fdfbf34e56
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Ora è il momento di impostare il nome del bot.

Crea una variabile chiamata `botName` e assegnale il valore stringa `"teacherBot"`.

# --hints--

Dovresti avere una variabile chiamata `botName`.

```js
assert.isNotNull(botName);
```

La tua variabile `botName` dovrebbe essere una stringa.

```js
assert.isString(botName);
```

La tua variabile `botName` dovrebbe avere il valore `"teacherBot"`.

```js
assert.strictEqual(botName, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

--fcc-editable-region--

--fcc-editable-region--
```
