---
id: 66b6e62423e8a031d6c1f03d
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Ora è il momento di ottenere la lunghezza della stringa `topic`.

Puoi usare i template literal all'interno delle istruzioni `console` in questo modo:

```js
const developer = "Jessica";
console.log(`Hello, my name is ${developer}.`);
```

Inizia mostrando il messaggio `Here is an example of using the length property on the word [topic].` nella console.

Ricorda di sostituire `[topic]` con la variabile `topic` e di usare la sintassi corretta dei template literal come hai fatto nei passaggi precedenti.

Poi, aggiungi una seconda istruzione `console.log` che mostri nella console la lunghezza della stringa `topic`.

# --hints--

Dovresti usare la sintassi dei template literal per mostrare il messaggio `Here is an example of using the length property on the word [topic].` nella `console`.

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{topic\}\.`\)/);
```

Dovresti avere una seconda istruzione `console` per mostrare la lunghezza della variabile `topic` nella `console`.

```js
assert.match(code, /console\.log\(\s*topic\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);
console.log(subject.length);

--fcc-editable-region--

--fcc-editable-region--
```
