---
id: 66b6fdb76441c738719039fa
title: Passaggio 13
challengeType: 1
dashedName: step-13
---

# --description--

Nelle lezioni precedenti, hai imparato come consultare l'ultimo carattere in una stringa in questo modo:

```js
const firstName = "Jessica";

// returns "a"
firstName[firstName.length - 1];
```

`string.length - 1` ti darà sempre l'ultimo numero di indice per una stringa.

Crea una nuova variabile chiamata `lastCharacter` e assegnale il valore dell'ultimo carattere nella variabile `subject`.

Poi, stampa il valore della variabile `lastCharacter` nella console.

# --hints--

Dovresti avere una variabile chiamata `lastCharacter`.

```js
assert.isNotNull(lastCharacter);
```

Dovresti assegnare il valore dell'ultimo carattere nella variabile `subject` alla variabile `lastCharacter`. Consulta l'esempio se hai bisogno di aiuto.

```js
assert.strictEqual(lastCharacter, subject[subject.length - 1]);
```

Dovresti stampare il valore della variabile `lastCharacter` nella `console`.

```js
assert.match(code, /console\.log\(lastCharacter\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
