---
id: 66b6f586767a1534f3097353
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Ora è il momento di consultare la seconda lettera della variabile `subject`.

Inizia aggiungendo un'istruzione `console` che produca in output il messaggio `Here is an example of accessing the second letter in the word [subject].` Ricorda di sostituire `[subject]` con il valore reale della variabile `subject` e di usare la sintassi corretta per i template literal.

Poi aggiungi un'altra istruzione `console` che produca in output la seconda lettera della variabile `subject` usando la notazione a parentesi e il numero di indice corretto.

# --hints--

La tua istruzione `console` dovrebbe produrre in output il messaggio `Here is an example of accessing the second letter in the word [subject].`

```js
assert.match(code, /console\.log\(`Here\s+is\s+an\s+example\s+of\s+accessing\s+the\s+second\s+letter\s+in\s+the\s+word\s+\$\{subject\}\.`\);?/);
```

Dovresti avere un'altra istruzione `console` che produca in output la seconda lettera della variabile `subject` usando la notazione a parentesi con indice `1`.

```js
assert.match(code, /console\.log\(subject\[1\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

--fcc-editable-region--

--fcc-editable-region--
```
