---
id: 66b6e39031393e30f2c48d0f
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Per ottenere la lunghezza di una stringa, puoi usare la proprietà <dfn>length</dfn>. Questa proprietà restituisce il numero di caratteri in una stringa.

Ecco un esempio:

```js
const greeting = "Hello, world!";

// Output: 13
console.log(greeting.length);
```

Inizia usando la proprietà `length` per ottenere la lunghezza della stringa `subject` e registra quel valore nella console.

# --hints--


Dovresti usare la proprietà `length` per ottenere la lunghezza della stringa `subject`.

```js
assert.match(code, /subject\.length/);
```

Dovresti registrare la lunghezza della stringa `subject` nella console.

```js
assert.match(code, /console\.log\(\s*subject\.length\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

--fcc-editable-region--

--fcc-editable-region--
```
