---
id: 66b6efddeca35833cd6f0b03
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Nelle lezioni precedenti, hai imparato come consultare i caratteri in una stringa in questo modo:

```js
const firstName = "Jessica";
// returns "J"
firstName[0];
```  

Ricorda che gli indici partono da `0`, quindi la prima lettera in una stringa sarà sempre all’indice `0`.

Inizia aggiungendo un’altra istruzione `console`.

All’interno dell’istruzione `console`, mostra la prima lettera della variabile `subject` usando la notazione a parentesi e l’indice corretto.

# --hints--

La tua istruzione `console` dovrebbe mostrare la prima lettera della variabile `subject` usando la notazione a parentesi e l’indice `0`.

```js
assert.match(code, /console\.log\(subject\[0\]\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

--fcc-editable-region--

--fcc-editable-region--
```
