---
id: 66b6ffb42b88e33943788abf
title: Passaggio 14
challengeType: 1
dashedName: step-14
---

# --description--

Per l'ultima parte del workshop, rivedrai come trovare la posizione indice di una sottostringa in una stringa. Ricorda che una sottostringa è una parte di una stringa.

Inizia creando una variabile chiamata `learningIsFunSentence` e assegnale il valore stringa `"Learning is fun."`.

# --hints--

Dovresti avere una variabile chiamata `learningIsFunSentence`.

```js
assert.isNotNull(learningIsFunSentence);
```

La tua `learningIsFunSentence` dovrebbe contenere il valore di una stringa.

```js
assert.isString(learningIsFunSentence);
```

Dovresti assegnare il valore di `"Learning is fun."` alla variabile `learningIsFunSentence`.

```js
assert.equal(learningIsFunSentence, "Learning is fun.");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

--fcc-editable-region--

--fcc-editable-region--
```
