---
id: 66b6d482bbb9e12f2e5ee1ae
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Per questa parte successiva del progetto, il bot insegnerĂ  a lavorare con la proprietĂ  string `length`.

Inizia creando una nuova variabile chiamata `strLengthIntro`.

Poi, usando la sintassi dei template literal, assegna la stringa `Here is an example of using the length property on the word [subject].` alla variabile `strLengthIntro`.

Sostituisci `[subject]` con la variabile `subject` come hai fatto prima.

Infine, stampa la variabile `strLengthIntro` nella console.

# --hints--

Dovresti avere una variabile chiamata `strLengthIntro`.

```js
assert.isNotNull(strLengthIntro);
```

La tua variabile `strLengthIntro` dovrebbe essere una stringa.

```js
assert.isString(strLengthIntro);
```

Dovresti usare la sintassi dei template literal per assegnare la stringa `Here is an example of using the length property on the word [subject goes here].` alla variabile `strLengthIntro`.

```js
assert.match(code,/`Here\s+is\s+an\s+example\s+of\s+using\s+the\s+length\s+property\s+on\s+the\s+word\s+\$\{subject\}\.`/);
```

Dovresti stampare la variabile `strLengthIntro` nella console.

```js
assert.match(code,/console\.log\(\s*strLengthIntro\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

--fcc-editable-region--

--fcc-editable-region--
```
