---
id: 66b7137348cfb53fd3ec6c73
title: Passaggio 17
challengeType: 1
dashedName: step-17
---

# --description--

Successivamente, aggiungi una nuova istruzione `console` che mostri il risultato dell'uso del metodo `indexOf` sulla variabile `learningIsFunSentence` per trovare la posizione della sottostringa `"fun"`.

Sotto quell'istruzione `console`, aggiungi una nuova istruzione `console` che mostri il risultato dell'uso del metodo `indexOf` per trovare la posizione della sottostringa `"learning"`.

Prendi nota di cosa mostra l'ultima istruzione `console`. 

# --hints--

La tua istruzione `console` dovrebbe mostrare la posizione indice della sottostringa `"fun"` dalla variabile `learningIsFunSentence`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(fun)\1\s*\)\s*\)/);
```

La tua seconda istruzione `console` dovrebbe mostrare la posizione indice della sottostringa `"learning"` dalla variabile `learningIsFunSentence`.

```js
assert.match(code, /console\.log\(\s*learningIsFunSentence\.indexOf\(\s*(['"])(learning)\1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");

const botName = "teacherBot";

const greeting = `My name is ${botName}.`;
console.log(greeting);

const subject = "JavaScript";
const topic = "strings";

const sentence = `Today, you will learn about ${topic} in ${subject}.`;
console.log(sentence);

const strLengthIntro = `Here is an example of using the length property on the word ${subject}.`;
console.log(strLengthIntro);

console.log(subject.length);

console.log(`Here is an example of using the length property on the word ${topic}.`);
console.log(topic.length);

console.log(`Here is an example of accessing the first letter in the word ${subject}.`);

console.log(subject[0]);

console.log(`Here is an example of accessing the second letter in the word ${subject}.`);
console.log(subject[1]);

console.log(`Here is an example of accessing the last letter in the word ${subject}.`);

const lastCharacter = subject[subject.length - 1];
console.log(lastCharacter);

const learningIsFunSentence = "Learning is fun.";

console.log("Here are examples of finding the positions of substrings in the sentence.");

console.log(learningIsFunSentence.indexOf("Learning"));

--fcc-editable-region--

--fcc-editable-region--
```
