---
id: 5d822fd413a79914d39e98d7
title: 第 15 步
challengeType: 0
dashedName: step-15
---

# --description--

要使用變數，請將變數名稱放在小括弧中，並在前面加上 `var`，像這樣：`var(--variable-name)`。你給變數的任何值都會套用到你使用它的任何屬性上。 

將你在前一步驟創建的變數 `--building-color1` 添加到 `.bb1a` 類別的 `background-color` 屬性值中。

# --hints--

應該設定 `.bb1a` 元素的 `background-color`。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

你應該使用 `var(--building-color1)` 來設定 `.bb1a` 元素的 `background-color`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}
--fcc-editable-region--
.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
