---
id: 5d822fd413a79914d39e98d1
title: 第 9 步
challengeType: 0
dashedName: step-9
---

# --description--

為你的 `.background-buildings` 元素設定 `width` 和 `height` 為 `100%`，使其成為父元素 `body` 的完整寬度和高度。

# --hints--

你應該使用 `background-buildings` 類別來選取正確的元素。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.background-buildings'));
```

你的 `.background-buildings` 元素應該有 `width` 為 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.width, '100%');
```

你的 `.background-buildings` 元素應該有 `height` 為 `100%`。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.background-buildings')?.height, '100%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings"></div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}
--fcc-editable-region--


--fcc-editable-region--
    
```

