---
id: 5d822fd413a79914d39e98c9
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

歡迎來到 CSS Variables Skyline 專案！請先在文件頂端添加 `!DOCTYPE html` 宣告，讓瀏覽器知道正在讀取何種類型的文件。

# --hints--

你的程式碼應該包含 `DOCTYPE` 引用。

```js
assert(code.match(/<!DOCTYPE/gi));
```

你應該在 `DOCTYPE` 引用後面加上一個空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

你應該定義文件型別為 `html`。

```js
assert(code.match(/html/gi));
```

你應該在 `DOCTYPE` 宣告式的型別後面以 `>` 關閉。

```js
assert(code.match(/html\s*>/gi));
```

你的 `DOCTYPE` 宣告應該放在你的 HTML 開頭。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
