---
id: 5d822fd413a79914d39e98cb
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

接著，為你的 `html` 元素添加開啟和關閉的 `head` 和 `body` 頁籤、標記。

# --hints--

你應該有一個開啟的 `head` 頁籤。

```js
assert(code.match(/<head\s*>/i));
```

你應該有一個結束的 `head` 頁籤。

```js
assert(code.match(/<\/head\s*>/i));
```

你應該有一個開啟的 `body` 頁籤。

```js
assert(code.match(/<body\s*>/i));
```

你應該有一個結束的 `body` 頁籤。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素應該是同層元素。

```js
assert(document.querySelector('head').nextElementSibling.localName === 'body');
```

`head` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'head'));
```

`body` 元素應該位於 `html` 元素內。

```js
assert([...document.querySelector('html').children].some(x => x.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html lang="en">
  
</html>
--fcc-editable-region--

```

