---
id: 5d822fd413a79914d39e98d6
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

現在你有一個類似建築物的東西。你已經準備好創建你的第一個變數。變數宣告以兩個破折號（`-`）開始，並給予名稱和值，如此：`--variable-name: value;`。在 `bb1` 類別的規則中，創建一個名為 `--building-color1` 的變數，並將其值設為 `#999`。

# --hints--

你應該創建一個名為 `--building-color1` 的新變數。

```js
assert(new __helpers.CSSHelp(document).isPropertyUsed('--building-color1'));
```

你應該在 `.bb1` 中定義 `--building-color1` 變數。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1'));
```

你應該給 `--building-color1` 一個 `#999` 的值。

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.getPropertyValue('--building-color1').trim(),'#999');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
--fcc-editable-region--
.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
}
--fcc-editable-region--
.bb1a {
  width: 70%;
  height: 10%;
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
