---
id: 6551fb531ce7f773071a0b84
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Como você pode ver, a saída de imprimir `type(text)` é `<class 'str'>`, o que significa que sua variável é uma string, indicada como `str`.

Agora vá para uma nova linha e crie outra variável chamada `shift` e atribua o valor `3` a essa variável.

# --hints--

Você deve declarar uma variável chamada `shift`. Preste atenção para colocar o nome da variável no início da linha.

```js
assert.match(code, /^shift\s*=/m)
```

Você deve atribuir o valor `3` à sua variável `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
