---
id: 6551f3bfed58796ebb1268dc
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Cada caractere da string pode ser referenciado por um índice numérico.A contagem do índice começa em zero.Então o primeiro caractere de uma string tem um índice de `0`.Por exemplo, na string `'Hello World'`, `'H'` está no índice `0`, `'e'` está no índice `1` e assim por diante. 

Cada caractere de uma string pode ser acessado usando a notação de colchetes.Você precisa escrever o nome da variável seguido por colchetes e adicionar o índice do caractere entre os colchetes:

```py
text = 'Hello World'
r = text[8]
```

Agora, em vez de imprimir `text`, imprima apenas o caractere no índice `6`.

# --hints--

Você ainda deve chamar a função `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Você deve passar `text[6]` para a função `print()` incluindo-o entre os parênteses. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
