---
id: 6551f22653e6ce6e4f1d1370
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Modifique sua chamada `print(shift)` para imprimir o tipo de dado da sua variável `shift`.

# --hints--

Você deve manter sua chamada existente `print(type(text))`.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Você deve passar `shift` para a função `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Você deve passar `type(shift)` para sua função `print()`. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

Você não deve ter `print(shift)` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
