---
id: 6551f1e8737dbe6e0cb08ea4
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

E agora imprima sua nova variável.

# --hints--

Você não deve remover sua chamada existente `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Você deve chamar a função `print()` passando a variável `shift`. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
