---
id: 6551f6f9ec6db270eb83f3db
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Você pode ver `11` impresso no terminal porque `'Hello World'` contém 11 caracteres.

Outra função interna útil é `type()`, que retorna o tipo de dado de uma variável. Modifique sua chamada `print()` para imprimir o tipo de dado de `text`.

# --hints--

Você deve chamar a função `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Você deve passar `text` para a função `type()` incluindo-o entre os parênteses.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Você deve passar `type(text)` para a função `print()` incluindo-o entre os parênteses. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Você não deve ter `print(len(text))` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
