---
id: 6551f5756c20146fc94f8675
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Você pode acessar o número de caracteres em uma string com a função embutida `len()`.

Modifique sua chamada existente de `print()` passando `len(text)` em vez de `text[-1]`.

# --hints--


Você deve chamar a função `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Você deve passar a variável `text` para a função `len()` incluindo-a entre os parênteses.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Você deve passar `len(text)` para a função `print()` incluindo-o entre os parênteses. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

Você não deve ter `print(text[-1])` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
