---
id: 6551ef504d91936d2d4e54f8
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Um *argumento* é um objeto ou uma expressão passada para uma função — adicionada entre os parênteses de abertura e fechamento — quando ela é chamada:

```py
greet = 'Hello!'
print(greet)
```

O código no exemplo acima imprimiria a string `'Hello!'`, que é o valor da variável `greet` passada para `print()` como argumento.

Imprima sua variável `text` na tela passando a variável `text` como argumento para a função `print()`.

# --hints--

Você deve passar `text` para a função `print()` adicionando o nome desta variável dentro dos parênteses de abertura e fechamento. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
