---
id: 6551f4259d66026ee28cb205
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Você também pode acessar caracteres de string começando pelo final da string.O último caractere tem um índice de `-1`, o penúltimo `-2` e assim por diante.

Agora modifique sua chamada existente de `print()` para imprimir o último caractere da sua string.

# --hints--

Você ainda deve chamar a função `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Você deve passar `text[-1]` para a função `print()` incluindo-o entre os parênteses. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
