---
id: 6551ecf597038c36637f3759
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Variáveis são uma parte essencial do Python e de qualquer linguagem de programação. Uma variável é um nome que referencia ou aponta para um objeto. Você pode declarar uma variável escrevendo o nome da variável no lado esquerdo do operador de atribuição `=` e especificando o valor a ser atribuído a essa variável no lado direito do operador de atribuição:

```py
variable_name = value
```

Crie uma variável chamada `number` e atribua o valor `5` à sua nova variável.

# --hints--

Você deve declarar uma variável chamada `number`. Preste atenção para colocar o nome da variável no início da linha.

```js
assert.match(code, /^number\s*=/m)
```

Você deve atribuir o valor `5` à sua variável `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
