---
id: 65521203d7165c7b84b22ad4
title: Passo 23
challengeType: 20
dashedName: step-23
---

# --description--

Um loop permite que você percorra sistematicamente uma sequência de elementos e execute ações em cada um.

Neste caso, você usará um loop `for`. Aqui está como você pode iterar sobre `text`:

```py
for i in text:
```

`for` é a palavra-chave que denota o tipo de loop. `i` é uma variável que assume sequencialmente o valor dos elementos em `text`. A declaração termina com dois pontos, `:`.

Abaixo da linha onde você declarou `alphabet`, escreva um loop `for` para iterar sobre `text`. Use `i` como a variável do loop.

Fazendo isso, há um erro no terminal. Você vai aprender sobre isso na próxima etapa.

# --hints--

Você deve usar a palavra-chave `for` para criar um loop. Certifique-se de colocar a palavra-chave `for` no início da linha e deixar um espaço em branco após a palavra-chave.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Você deve escrever a variável `i` após a palavra-chave `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Você deve escrever a palavra-chave `in` após `for i`. Certifique-se de deixar um espaço ao redor da palavra-chave `in`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Você deve escrever `text` após `for i in ` . Não esqueça de adicionar o `:` final.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Seu loop `for` deve ser colocado abaixo da linha de código `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
