---
id: 655261b2e1f2c197093f3993
title: Passo 41
challengeType: 20
dashedName: step-41
---

# --description--

Atualmente, os espaços são criptografados como `'c'`. Para manter o espaçamento original na mensagem simples, você precisará de uma declaração condicional `if`. Isto é composto pela palavra-chave `if`, uma condição e dois-pontos `:`.

```py
if x != 0:
    print(x)
```

No exemplo acima, a condição da declaração `if` é `x != 0`. O código `print(x)`, dentro do corpo da instrução `if`, é executado somente quando a condição avalia para `True` (neste exemplo, significando que `x` é diferente de zero).

At the top of your for loop, replace `print(char == ' ')` with an `if` statement.  The condition of this `if` statement should evaluate to `True` if `char` is an empty space and `False` otherwise. Dentro do corpo do `if`, imprima a string `'space!'`. Lembre-se de indentar esta linha.

# --hints--

Você não deve ter `print(char == ' ')` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Você deve substituir `print(char == ' ')` por uma declaração `if` que seja acionada quando `char == ' '`. Não use parênteses para envolver a condição `if` e lembre-se de incluir os dois pontos finais.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Você deve imprimir a string `'space!'` dentro do seu novo `if` statement.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
