---
id: 655220953ba90d80514d7ee2
title: Passo 27
challengeType: 20
dashedName: step-27
---

# --description--

Atualmente, a função `print()` está recebendo um único argumento `char`, mas ela pode receber múltiplos argumentos, separados por vírgula.

Adicione um segundo argumento a `print(char)` para que ele imprima o caractere e seu índice dentro do alfabeto.

# --hints--

Você deve adicionar `index` como o segundo argumento na sua chamada existente `print(char)`. Não esqueça de separar os argumentos com uma vírgula.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*,\s*index\s*\)\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char)
--fcc-editable-region--
```
