---
id: 655208d59b131e7816f18c96
title: Passo 15
challengeType: 20
dashedName: step-15
---

# --description--

A função `print()` fornece apenas uma saída no console, mas funções e métodos podem ter um valor de retorno que você pode usar no seu código.

Agora atribua `alphabet.find(text[0])` a uma variável chamada `index`.Dessa forma, `index` armazenará o valor retornado por `alphabet.find(text[0])`.

# --hints--

Você deve declarar uma variável chamada `index`. Preste atenção para colocar o nome da variável no início da linha.

```js
assert.match(code, /^index\s*=/m)
```

Você deve atribuir o valor retornado por `alphabet.find(text[0])` à sua variável `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
