---
id: 65520e6f2b9678799977f24d
title: Passo 18
challengeType: 20
dashedName: step-18
---

# --description--

Remova a última chamada `print()`.Então, em vez de `text[0]`, passe `text[0].lower()` como argumento para sua chamada `.find()` e veja a saída.

# --hints--

Você não deve ter `print(text.lower())` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Você ainda deve imprimir sua variável `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Você deve atualizar sua chamada `alphabet.find(text[0])` para usar `text[0].lower()` como argumento. Preste atenção para colocar a chamada do método no início da linha.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
