---
id: 65525e359ca28d938baa82c5
title: Passo 39
challengeType: 20
dashedName: step-39
---

# --description--

Você pode obter o mesmo efeito de `a = a + b` usando o operador de atribuição de adição:

```py
a += b
```

O operador de atribuição de adição permite que você some um valor a uma variável e depois atribua o resultado a essa variável.

Use o operador `+=` para adicionar um valor e atribuí-lo ao mesmo tempo a `encrypted_text`.

# --hints--

Você deve usar o operador de atribuição de adição para adicionar `alphabet[new_index]` ao valor atual de `encrypted_text` dentro do corpo do seu loop.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]/));
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''
--fcc-editable-region--
for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text = encrypted_text + alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
