---
id: 65522255d5b9cd80f335c6f2
title: Passo 29
challengeType: 20
dashedName: step-29
---

# --description--

No final do corpo do seu loop, declare uma variável chamada `new_index` e atribua o valor de `index + shift` a essa variável.

# --hints--

Você deve declarar uma variável chamada `new_index` dentro do seu loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Você deve atribuir `index + shift` à sua nova variável no final do corpo do seu loop `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
