---
id: 65521fc818947e800bffe48a
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

Dentro do loop `for`, antes de imprimir o caractere atual, declare uma variável chamada `index` e atribua a essa variável o valor retornado por `alphabet.find(char)`.

# --hints--

Você deve declarar uma nova variável chamada `index` no início do seu loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Você deve atribuir `alphabet.find(char)` à sua nova variável `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
