---
id: 655251308f31958d06cdf267
title: Passo 36
challengeType: 20
dashedName: step-36
---

# --description--

No momento, o caractere criptografado é atualizado a cada iteração. Seria melhor armazenar a string criptografada em uma nova variável. Antes do seu loop `for`, declare uma variável chamada `encrypted_text` e atribua uma string vazia (`''`) a essa variável.

# --hints--

Você deve declarar uma variável chamada `encrypted_text` antes do seu loop `for`.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Você deve atribuir uma string vazia à sua variável `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
