---
id: 65520f8282faf57a0db4f7fe
title: Passo 21
challengeType: 20
dashedName: step-21
---

# --description--

Como você pode ver na saída, `'h'` está no índice `7` na string `alphabet`.Agora você precisa encontrar a letra no índice `7` mais o valor de `shift`.Para isso, você pode usar o operador de adição, `+`, da mesma forma que usaria para uma adição matemática.

Modifique sua variável `shifted` para que ela armazene o valor de `alphabet` no índice `index + shift`.

# --hints--

Você deve atribuir `alphabet[index + shift]` à sua variável `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
