---
id: 6552008c0d9d9075cbec9772
title: Passo 13
challengeType: 20
dashedName: step-13
---

# --description--

Você vai usar o método `.find()` para encontrar a posição no alfabeto de cada letra na sua mensagem. Um método é semelhante a uma função, mas pertence a um objeto.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Acima, o método `.find()` é *chamado em* `sentence` (a string para buscar), e `'r'` (o caractere a localizar) é passado como argumento. A chamada `sentence.find('r')` retornará `4`, que é o índice da primeira ocorrência de `'r'` em `sentence`.

No final do seu código, chame `.find()` em `alphabet` e passe `'z'` como argumento para o método.

# --hints--

Você deve chamar o método `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Você deve chamar o método `.find()` na variável `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Você deve chamar o método `.find()` na variável `alphabet` e passar `'z'` para o método. Preste atenção para colocar a chamada do método no início da linha.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
