---
id: 655209a4a27dd37873c4cac3
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` retorna o índice do caractere correspondente dentro da string. Se o caractere não for encontrado, ele retorna `-1`. Como você pode ver, o primeiro caractere em `text`, maiúsculo `'H'`, não é encontrado, já que `alphabet` contém apenas letras minúsculas.

Você pode transformar uma string em seu equivalente em letras minúsculas com o método `.lower()`. Adicione outra chamada `print()` para imprimir `text.lower()` e veja a saída.

# --hints--

Você ainda deve ter `print(index)` no seu código. Preste atenção para ter a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Você deve imprimir `text.lower()`. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
