---
id: 65524d2a1a253b8bb5197ae2
title: Passo 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` está novamente retornando `-1` para letras maiúsculas e para o caractere de espaço também.Você vai cuidar do espaço mais tarde.

Por enquanto, em vez de iterar sobre `text`, altere o loop `for` para iterar sobre `text.lower()`.

# --hints--

Você deve modificar seu loop `for` para iterar sobre `text.lower()` em vez de `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
