---
id: 6552487e689f6e895f658717
title: Passo 33
challengeType: 20
dashedName: step-33
---

# --description--

Agora você precisa criar uma variável `new_char` no final do corpo do seu loop. Defina seu valor para `alphabet[new_index]`.

# --hints--

Você deve criar uma variável `new_char` dentro do seu loop `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Você deve definir sua variável `new_char` para `alphabet[new_index]` no final do corpo do seu loop.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
