---
id: 6553980e0527fa115c705646
title: Passo 42
challengeType: 20
dashedName: step-42
---

# --description--

Agora, em vez de imprimir `'space!'`, use o operador de atribuição de adição para adicionar o espaço (atualmente armazenado em `char`) ao valor atual de `encrypted_text`.

# --hints--

Você não deve ter `print('space!')` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*("|')space!\1\s*\)/);
```

Você deve usar o operador `+=` para adicionar `char` ao valor atual de `encrypted_text` dentro do seu novo comando `if`.

```js
const commentless_code = __helpers.python.removeComments(code);    
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);
assert(block_body.match(/encrypted_text\s*\+=\s*char/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        print('space!')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
