---
id: 6553a572f7a65718f1e42e18
title: Passo 49
challengeType: 20
dashedName: step-49
---

# --description--

Uma função é essencialmente um bloco reutilizável de código. Você já conheceu algumas funções internas, como `print()`, `find()` e `len()`. Mas você também pode definir funções personalizadas assim:

```py
def function_name():
    <code>
```

Uma declaração de função começa com a palavra-chave `def` seguida pelo nome da função — um nome de variável válido — e um par de parênteses. A declaração termina com dois pontos.

Logo após a sua variável `shift`, declare uma função chamada `caesar` e indente todas as linhas seguintes para dar um corpo à sua nova função.

# --hints--

Você deve usar a palavra-chave `def` para declarar uma nova função.

```js
assert.match(code, /^def\s+/m)
```

Você deve escrever `caesar` como nome da função após a palavra-chave `def`. Lembre-se de adicionar um espaço após `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Você deve adicionar um par de parênteses após o nome da função. Não esqueça do dois-pontos final.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Você deve indentar todas as linhas após `shift = 3` para que elas se tornem o novo corpo da sua função.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
