---
id: 6553f6b4ec51112d44d737c8
title: Passo 48
challengeType: 20
dashedName: step-48
---

# --description--

Logo antes da chamada `print`, adicione outra e passe `'plain text:', text` como argumentos para `print()`.Use a mesma indentação.

# --hints--

Você deve imprimir `'plain text:', text` após seu loop for.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*^print/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
