---
id: 6553a7d8d05cbb1ae335a665
title: Passo 45
challengeType: 20
dashedName: step-45
---

# --description--

Quando o loop alcança a letra `Z`, a soma `index + shift` excede o último índice da string `alphabet`.Portanto, `alphabet[new_index]` está tentando usar um índice inválido, o que causa o lançamento de um `IndexError`.

Você pode notar que a saída no terminal para no espaço imediatamente antes do `Z`, o último `print` antes do erro ser lançado.

Neste caso, o operador modulo (`%`) pode ser usado para retornar o resto da divisão entre dois números.Por exemplo: `5 % 2` é igual a `1`, porque 5 dividido por 2 tem um quociente de 2 e um resto de 1.

Coloque `index + shift` entre parênteses e faça o módulo da expressão com `26`, que é o comprimento do `alphabet`.

# --hints--

Você deve ter `new_index = (index + shift) % 26` na sua declaração `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
