---
id: 6553efd6ada3f42aa2d75448
title: Passo 53
challengeType: 20
dashedName: step-53
---

# --description--

Agora você deve ver a saída novamente.Embora essa abordagem funcione, ela não melhora significativamente a reutilização do código.Chamar sua função repetidamente resultará no mesmo resultado.No entanto, funções podem ser declaradas com *parâmetros* para introduzir versatilidade e personalização:

```py
def function_name(param_1, param_2):
    <code>
```

Parâmetros são variáveis que você pode usar dentro da sua função.Uma função pode ser declarada com diferentes números de parâmetros.No exemplo acima, `param_1` e `param_2` são parâmetros.

Modifique a declaração da sua função para que ela receba dois parâmetros chamados `message` e `offset`.

Depois disso, você verá um erro aparecer no terminal.Você verá como resolver isso nos próximos passos.

# --hints--

Sua função `caesar` deve receber `message` e `offset` como parâmetros.Lembre-se de separar os parâmetros com uma vírgula.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
