---
id: 6553f4f66099802c6ae94613
title: Passo 55
challengeType: 20
dashedName: step-55
---

# --description--

Atualmente, seu código gera um `TypeError`, porque a função `caesar` é definida com dois parâmetros (`message` e `offset`), portanto ela espera ser chamada com dois *argumentos*.

Chamar `caesar()` sem os argumentos necessários interrompe a execução do código.

Forneça os valores `message` e `offset`, passando `text` e `shift` como argumentos para a chamada da função `caesar`.

# --hints--

Você deve passar `text` e `shift` como argumentos para a chamada da sua função incluindo-os dentro dos parênteses. Não esqueça de separar os argumentos com uma vírgula.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
