---
id: 6553995f412dd8122ed38e4a
title: Passo 43
challengeType: 20
dashedName: step-43
---

# --description--

Uma declaração condicional também pode ter uma cláusula `else`.Esta cláusula pode ser adicionada ao final de uma declaração `if` para executar um código alternativo se a condição da declaração `if` for falsa:

```py
if x != 0:
    print(x)
else:
    print('x = 0')
```

Como você pode ver na sua saída, quando as iterações do loop alcançam o espaço, um espaço é adicionado à string criptografada.Então o código fora do bloco `if` é executado e um `c` é adicionado à string criptografada.

Para corrigir, adicione uma cláusula `else` após `encrypted_text += char` e indente todas as linhas de código subsequentes exceto a chamada `print()`.

# --hints--

Você deve criar uma cláusula `else`. Lembre-se de incluir os dois pontos finais.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/else\s*:/));
```

Você deve indentar as linhas de código após sua cláusula `else` exceto a chamada `print()`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*new_index\s*=\s*index\s*\+\s*shift\s*encrypted_text\s*\+=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
