---
id: 6553f6086add4b2cbb99fd78
title: Passo 56
challengeType: 20
dashedName: step-56
---

# --description--

Na parte inferior do seu código, após sua chamada existente `caesar(text, shift)`, chame sua função novamente.

Desta vez, passe a variável `text` e o inteiro `13` como argumentos.

# --hints--

Você deve chamar sua função novamente, desta vez passando `text` e `13` como argumentos.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*13\s*\)/m) })
```

Você deve manter a chamada de função existente.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*text\s*,\s*shift\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
--fcc-editable-region--
```
