---
id: 6553ed69ece88d29594748aa
title: Passo 52
challengeType: 20
dashedName: step-52
---

# --description--

Para executar, uma função precisa ser chamada (ou invocada) adicionando um par de parênteses após seu nome, assim:

```py
function_name()
```

No final do seu código, chame sua função `caesar`.Preste atenção na indentação.

# --hints--

Você deve chamar sua função `caesar`.Certifique-se de escrever a chamada da função no início da linha.

```js
({ test: () => assert.match(code, /^caesar\s*\(\s*\)/m) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
