---
id: 6553f8c570f9982e013a8886
title: Passo 57
challengeType: 20
dashedName: step-57
---

# --description--

Atualmente, cada letra é sempre criptografada com a mesma letra, dependendo do `offset` especificado. E se o deslocamento fosse diferente para cada letra? Isso seria muito mais difícil de descriptografar. Este algoritmo é conhecido como cifra de Vigenère, onde o deslocamento para cada letra é determinado por outro texto, chamado de *key*.

Comece a transformar seu cifrador de César em um cifrador de Vigenère removendo as duas chamadas de função.

# --hints--

Você deve remover suas duas chamadas da função `caesar()`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*shift\s*\)/m.test(commentless_code))
    assert.isFalse(/^caesar\s*\(\s*text\s*,\s*13\s*\)/m.test(commentless_code))
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

caesar(text, shift)
caesar(text, 13)
--fcc-editable-region--
```
