---
id: 6553a755879b131a445e664c
title: Passo 44
challengeType: 20
dashedName: step-44
---

# --description--

Tente atribuir a string `'Hello Zaira'` à sua variável `text` e veja o que acontece no terminal.

Você verá uma exceção de `string index out of range`.Não se preocupe, você vai descobrir como consertar isso em breve!

# --hints--

Você deve ter uma variável `text`.

```js
assert.match(code, /^text\s*=/m)
```

Você deve atribuir a string `'Hello Zaira'` à sua variável `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
