---
id: 655491bd5b98b813fa5bedca
title: Passo 58
challengeType: 20
dashedName: step-58
---

# --description--

Agora modifique a declaração da sua função: mude o nome da função para `vigenere` e o segundo parâmetro para `key`.

# --hints--

Você deve modificar o nome da sua função para `vigenere`.

```js
assert.match(code, /^def\s+vigenere\s*\(.*\)\s*:/m)
```

Sua função `vigenere` deve receber `message` e `key` como parâmetros.

```js
assert.match(code, /^def\s+vigenere\s*\(\s*message\s*,\s*key\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
