---
id: 6554d25dc5ceaa354307a77e
title: Passo 70
challengeType: 20
dashedName: step-70
---

# --description--

Criptografia e descriptografia são processos opostos e sua função pode fazer ambos com algumas modificações.

Adicione um terceiro parâmetro chamado `direction` à definição da sua função. Além disso, comente as duas últimas linhas de código para evitar erros no console.

# --hints--

Você deve transformar as duas últimas linhas do seu código em comentários. Coloque um `#` no início de cada linha.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

Você deve adicionar `direction` como o terceiro parâmetro da sua função.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
