---
id: 6554ad2463b8892748f8efdd
title: Passo 67
challengeType: 20
dashedName: step-67
---

# --description--

No momento, sua função imprime algumas strings, mas esses valores não podem ser usados por outras partes do código para executar quaisquer ações.

Para isso, você precisa usar uma declaração `return`:

```python
def foo():
    return 'spam'
```

Você precisa escrever `return` seguido de um espaço e do valor que a função deve retornar.Uma vez que a instrução `return` é encontrada, esse valor é retornado e a execução da função para, prosseguindo para a próxima linha de código após a chamada da função.No exemplo acima, a função `foo` retorna a string `'spam'`.

Remova as duas chamadas `print()` da sua função e retorne `encrypted_text`.

# --hints--

Você deve remover as duas chamadas `print()` da sua função.

```js
({ 
  test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.isFalse(/print\s*\(\s*("|')plain\stext:\1\s*,\s*message\s*\)/.test(commentless_code))
    assert.isFalse(/print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/.test(commentless_code))
  }
})
```

Sua função deve retornar `encrypted_text`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+message\.lower\s*\(\s*\)\s*/);
    const regex = /return\s+encrypted_text/
    assert(function_body.match(regex));
    assert.notMatch(block_body, regex);
  }
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
--fcc-editable-region--
            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
