---
id: 6554930320d70414e7b6acc6
title: Passo 59
challengeType: 20
dashedName: step-59
---

# --description--

Delete sua variável `shift`.Então, declare uma nova variável chamada `custom_key` e atribua o valor da string `'python'` a essa variável.

# --hints--

Você deve deletar a variável `shift` e seu valor.

```js
({ test: () => assert.isFalse(__userGlobals.has("shift")) })
```

Você deve declarar uma variável chamada `custom_key`.

```js
({ test: () => assert(__userGlobals.has("custom_key")) })
```

Você deve atribuir a string `'python'` à sua variável `custom_key`.

```js
({ test: () => assert.equal(__userGlobals.get("custom_key"), "python") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
