---
id: 6554a88d5af937226f4a9121
title: Passo 65
challengeType: 20
dashedName: step-65
---

# --description--

O método `.index()` é idêntico ao método `.find()` mas lança uma exceção `ValueError` se não conseguir encontrar a substring.

Um `ValueError` é uma exceção embutida que é levantada quando um argumento com o tipo correto mas valor inadequado é passado para uma função.

Após incrementar `key_index`, declare uma variável chamada `offset`. Encontre o índice que `key_char` tem em `alphabet` e atribua-o a `offset`. Use o `.index()` para encontrar o índice.

# --hints--

Você deve declarar uma variável chamada `offset`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=/));
```

Sua variável `offset` deve armazenar o valor de `alphabet.index(key_char)`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/offset\s*=\s*alphabet\.index\s*\(\s*key_char\s*\)/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--
        else:
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1
            
--fcc-editable-region--            
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
