---
id: 65549f90cf78131c96ebcf28
title: Passo 61
challengeType: 20
dashedName: step-61
---

# --description--

Ao programar, a legibilidade é fundamental.Comentários servem como notas eficazes, explicando a lógica por trás do seu código.Eles são valiosos ao retornar a um projeto após algum tempo e também ajudam seus colegas a entender o código.

Em Python, você pode escrever um comentário usando um `#`.Qualquer coisa que vier depois do `#` não será executada.

Antes do seu `if` statement, adicione um comentário dizendo `Append space to the message`.

# --hints--

Você deve adicionar um comentário dizendo `Append space to the message`.Não esqueça do caractere `#` no início.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
