---
id: 6555e43e783ed31a0532b1b2
title: Passo 93
challengeType: 20
dashedName: step-93
---

# --description--

O caractere de nova linha `\n` é uma sequência especial usada para representar uma nova linha. Você pode escrever uma barra invertida `\` seguida de um `n` como uma sequência normal de caracteres em uma string e ela será substituída por uma nova linha na saída quando o programa for executado.

Coloque um caractere de nova linha no início da sua primeira chamada `print` e veja a saída.

# --hints--

Você deve modificar sua primeira chamada de print para `print(f'\nEncrypted text: {text}')`.

```js
({ test: () => assert.match(code, /^print\s*\(\s*f("|')\\nEncrypted\stext:\s\{\s*text\s*\}\1\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print(f'Encrypted text: {text}')
print(f'Key: {custom_key}')
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
