---
id: 6555d6e12c95701172b55709
title: Passo 77
challengeType: 20
dashedName: step-77
---

# --description--

Funções podem ser chamadas com argumentos padrão. Um argumento padrão indica o valor que a função deve assumir se o argumento não for passado. Por exemplo, a função abaixo aceita três argumentos mas você pode chamá-la passando dois argumentos. O terceiro assumirá o valor especificado por padrão:

```py
def foo(a, b, c=value):
    <code>
```

Modifique a função `vigenere` para que seu parâmetro `direction` tenha um valor padrão de `1`.

# --hints--

O parâmetro `direction` da sua função `vigenere` deve ter um valor padrão de `1`.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction=1)'
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
custom_key = 'python'

--fcc-editable-region--
def vigenere(message, key, direction):
--fcc-editable-region--
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
    
encryption = vigenere(text, custom_key, 1)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
```
