---
id: 6555231eac4d9f19bd3d44b5
title: Passo 83
challengeType: 20
dashedName: step-83
---

# --description--

A palavra-chave `pass` pode ser usada como um espaço reservado para código futuro.Não tem nenhum efeito no seu código mas pode te salvar de erros que você teria em caso de código incompleto:

```py
def foo():
    pass
```

Chamar `vigenere` com `1` para criptografar e `-1` para descriptografar é aceitável mas pode ser um pouco criptográfico.Crie uma nova função chamada `encrypt` que recebe os parâmetros `message` e `key` e use `pass` para preencher o corpo da função.


# --hints--

Você deve declarar uma nova função chamada `encrypt`.Lembre-se de usar a palavra-chave `pass` dentro do corpo.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Sua função `encrypt` deve receber `message` e `key` como parâmetros.A ordem importa.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
