---
id: 6555e1bfcd374e18c6be8e58
title: Passo 91
challengeType: 20
dashedName: step-91
---

# --description--

Em Python, existe uma maneira fácil de formatar strings. *f-strings* permitem que você interpolar valores nas suas strings. 

Interpolação significa escrever espaços reservados que serão substituídos pelos valores especificados quando o programa for executado. Por exemplo, você pode obter o mesmo resultado de `'Encrypted text: ' + text` com `f'Encrypted text: {text}'`. Você precisa colocar um `f` antes das aspas para criar a f-string e escrever as variáveis ou expressões que deseja interpolar entre chaves.

Modifique a primeira chamada `print()` para usar uma f-string.

# --hints--

Você deve modificar a chamada `print('Encrypted text: ' + text)` para que ela imprima o mesmo resultado usando uma f-string.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
