---
id: 655a2a7210094920069b117c
title: Passo 40
challengeType: 20
dashedName: step-40
---

# --description--

Operadores de comparação permitem que você compare dois objetos com base em seus valores.Você pode usar um operador de comparação colocando-o entre os objetos que deseja comparar.Eles retornam um valor *Boolean* — ou seja, `True` ou `False` — dependendo da veracidade da expressão.

Python possui os seguintes operadores de comparação:

<table>
  <thead>
    <tr>
      <th>Operator</th>
      <th>Meaning</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Equal</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Not equal</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Greater than</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Less than</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Greater than or equal to</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Less than or equal to</td>
    </tr>
  </tbody>
</table>

No início do corpo do seu loop, imprima o resultado da comparação de `char` com um espaço (`' '`).Use o operador de igualdade `==` para isso.

# --hints--

Você deve comparar `char` com um espaço usando o operador de igualdade dentro do seu loop `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Você deve imprimir o resultado de comparar `char` com um espaço dentro do seu loop `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Você deve imprimir o resultado de comparar `char` com um espaço no início do seu loop.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
