---
id: 65688a50e6c998a21d8e41d3
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

Strings são imutáveis, o que significa que elas não podem ser alteradas depois de criadas.Por exemplo, você pode pensar que o seguinte código altera o valor de `my_string` para a string `'train'`, mas isso não é válido:

```py
my_string = 'brain'
my_string[0] = 't'
```

Confirme que, ao usar a notação de colchetes para acessar a primeira letra em `text`, você tenta alterá-la para um caractere de sua escolha.Você verá a saída desaparecer e um erro aparecer.

# --hints--

Você ainda deve ter `text = 'Hello World'` no seu código.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Você deve acessar a primeira letra em `text` com `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Você deve usar o operador `=` para atribuir um caractere de sua escolha a `text[0]`.Não esqueça de colocar o caractere entre aspas simples ou duplas.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
