---
id: 656886ec4e4d9b96ded3c565
title: Passo 25
challengeType: 20
dashedName: step-25
---

# --description--

A variável de iteração pode ter qualquer nome válido, mas é conveniente dar a ela um nome significativo.

Renomeie sua variável `i` para `char`.

# --hints--

Você deve substituir `i` por `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
