---
id: 65688b5a1655a7a6caede847
title: Passo 31
challengeType: 20
dashedName: step-31
---

# --description--

Quando você tenta alterar os caracteres individuais de uma string como fez na etapa anterior, você recebe um `TypeError`, que ocorre quando um objeto de tipo inadequado é usado no seu código.

Como você pode ver na mensagem de erro, strings não suportam atribuição de item porque elas são imutáveis.No entanto, uma variável pode ser reatribuída a outra string:

```py
message = 'Hello World'
message = 'Hello there!'
```

Delete a linha `text[0]` e reatribua a `text` a string `'Albatross'`.

# --hints--

Você não deve ter `text[0] = 'C'` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Você deve reatribuir `text` a string `'Albatross'`.Não modifique `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
