---
id: 65687d2f8c7ee27b0446cef3
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

O código a ser executado em cada iteração — colocado após o `:` — constitui o corpo do loop.Este código deve ser indentado.Em Python, é recomendado usar 4 espaços por nível de indentação.Este nível indentado é um bloco de código.

```py
for i in text:
    print(i)
```

Python depende da indentação para indicar blocos de código.Dois-pontos no final de uma linha é um sinal de que um novo bloco de código indentado seguirá.

Então, quando nenhum bloco indentado é encontrado após os dois pontos finais, a execução do código para e um `IndentationError` é lançado.Este código não exibirá a saída e em vez disso gerará um `IndentationError`:

```py
for i in text:
print(i)
```

Dê ao seu loop `for` um corpo adicionando uma chamada para `print(i)`.Lembre-se de indentar o corpo do loop.

# --hints--

Você deve adicionar `print(i)` ao corpo do seu loop `for`.Preste atenção na indentação. 

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
