---
id: 659530ca823d6c47f61daa8a
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Imprima sua variável `shifted`.

# --hints--

Você deve imprimir sua variável `shifted`. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
