---
id: 65951b637ddec63611c482c0
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

Você pode usar a função embutida `print()` para imprimir a saída do seu código no terminal.

Funções são blocos de código reutilizáveis que você pode chamar, ou invocar, para executar seu código quando precisar delas. Para chamar uma função, você só precisa escrever um par de parênteses ao lado do seu nome. Você aprenderá mais sobre funções muito em breve.

Por enquanto, vá para uma nova linha e adicione uma chamada vazia para a função `print()`. Você não deve ver nenhuma saída ainda.

# --hints--

Você deve ter `print()` no seu código. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
