---
id: 659526f3e5424b3eac4b6aa7
title: Passo 16
challengeType: 20
dashedName: step-16
---

# --description--

Em seguida, imprima a variável `index` no console.

# --hints--

Você deve imprimir sua variável `index`. Preste atenção para colocar a chamada da função no início da linha.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
