---
id: 65952b70924edb449faf1950
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

Declare uma nova variável chamada `shifted`. Use a notação de colchetes para acessar o valor de `alphabet` no índice `index` e atribuí-lo à sua nova variável.

# --hints--

Você deve declarar uma variável chamada `shifted`. Preste atenção para colocar o nome da variável no início da linha.

```js
assert.match(code, /^shifted\s*=/m)
```

Você deve atribuir `alphabet[index]` à sua variável `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
