---
id: 659db3e3670d3e712be82593
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Repetir o processo de localizar a letra dentro do alfabeto e determinar a letra deslocada para cada caractere em `text` pode ser tedioso.Felizmente, você pode simplificá-lo usando um loop.

Por enquanto, remova todas as linhas de código abaixo da declaração da variável `alphabet`.

# --hints--

Você ainda deve ter `text = 'Hello World'` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

Você ainda deve ter `shift = 3` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

Você ainda deve ter `alphabet = 'abcdefghijklmnopqrstuvwxyz'` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Você deve deletar a variável `index` e seu valor.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

Você não deve ter `print(index)` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Você deve deletar a variável `shifted` e seu valor.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

Você não deve ter `print(shifted)` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
