---
id: 65a450e8fb2c9d75c7378d28
title: Passo 88
challengeType: 20
dashedName: step-88
---

# --description--

Como desta vez você está começando a partir de uma string criptografada para descriptografar, você não precisará mais da variável `encryption`.

Delete `encryption` e a chamada `print(encryption)`.Além disso, comente as duas últimas linhas do seu código.

# --hints--

Você deve deletar a variável `encryption` e seu valor.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/encryption\s*=\s*encrypt\s*\(\s*text\s*,\s*custom_key\s*\)/.test(commentless_code))
```

Você não deve ter `print(encryption)` no seu código.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse(/print\s*\(\s*encryption\s*\)/.test(commentless_code))
```

Você deve transformar `decryption = decrypt(encryption, custom_key)` em um comentário.

```js
assert(code.match(/#\s*decryption\s*=/));
```

Você deve transformar `print(decryption)` em um comentário.

```js
assert(code.match(/#\s*print\s*\(\s*decryption\s*\)/));
```

Seu código contém sintaxe inválida ou indentação inválida.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
encryption = encrypt(text, custom_key)
print(encryption)
decryption = decrypt(encryption, custom_key)
print(decryption)
--fcc-editable-region--
```
