---
id: 6624e0ec0bf47017eed84b1f
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

O primeiro tipo de cifra que você vai construir é chamado de cifra *Caesar*. Especificamente, você vai pegar cada letra da sua mensagem, encontrar sua posição no alfabeto, pegar a letra localizada 3 posições depois no alfabeto e substituir a letra original pela nova letra.

Para implementar isso, você usará o método `.find()` discutido na etapa anterior. Modifique sua chamada `.find()` existente passando `text[0]` como argumento em vez de `'z'`.

# --hints--

Você deve modificar sua chamada existente `alphabet.find('z')` passando `text[0]` para o método.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
