---
id: 66605a64a9822f01ec5ae8ce
title: Passo 50
challengeType: 20
dashedName: step-50
---

# --description--

Em Python, o *escopo* de uma variável determina onde essa variável pode ser acessada:

- Variáveis definidas fora de uma função têm escopo *global* e podem ser acessadas de qualquer parte do seu código.

- Variáveis definidas dentro de uma função são *locais* a essa função e não podem ser acessadas fora dela.

Para ver isso em ação, tente imprimir a variável `alphabet` no final do seu código. Isso gerará uma exceção `NameError`.

Você deve ver uma mensagem de erro indicando que `alphabet` não está definido. Isso ocorre porque `alphabet` é definido dentro da função `caesar` e não é acessível fora dela.

# --hints--

Você deve tentar imprimir a variável `alphabet` fora da função caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
