---
id: 683ecb010f4c6bafd428802c
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、一連の医療データが一連のルールに準拠しているかどうかを検証します。

`medical_records` という名前の変数を宣言し、空のリストを代入してください。次のステップで、この変数を使って医療データをストアします。

# --hints--

`medical_records` という名前の変数を宣言してください。

```js
({ test: () => runPython(`
assert _Node(_code).has_variable("medical_records")
`) })
```

`medical_records` 変数に空のリストを代入してください。

```js
({ test: () => runPython(`
assert medical_records == []
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
