---
id: 6842aee5bcaf4ab8883f8191
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

前のレッスンで学んだように、辞書はキーと値のペアを保持するデータ構造で、キーは不変のデータ型（文字列など）であり、辞書内で一意です。

```py
person = {
    'name': 'John',
    'age': 33
}
```

`medical_records` リストは辞書を格納し、それぞれが患者を表します。`patient_id` リストに、キーが `P1001` で値が文字列の `medical_records` の辞書を追加してください。

# --hints--

`medical_records` リストには1つの辞書が含まれているはずです。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

辞書には、文字列 `patient_id` の値を持つ `P1001` キーがあるはずです。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
--fcc-editable-region--
    
--fcc-editable-region--
]
```
