---
id: 6842d5b537c64266420faf62
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

辞書に、整数 `age` の値を持つキー `34` と、文字列 `gender` の値を持つキー `Female` を追加してください。キーと値のペアの間のカンマを忘れないでください。

# --hints--

あなたの `medical_records` リストには、1つの辞書が含まれているはずです。

```js
({ test: () => runPython(`
assert isinstance(medical_records[0], dict) and len(medical_records) == 1
`) })
```

辞書には、整数 `age` の値を持つ `34` キーがあるはずです。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

辞書には、文字列 `gender` の値を持つ `Female` キーがあるはずです。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

辞書には、文字列 `patient_id` の値を持つ `P1001` キーがあるはずです。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
--fcc-editable-region--
        'patient_id': 'P1001'
        
--fcc-editable-region--
    }    
]
```
