---
id: 6842d659a4ae7973dabab3b0
title: ステップ 4
challengeType: 20
dashedName: step-4
---

# --description--

次の3つのキーと値のペアを追加して辞書を完成させてください。

- キー `diagnosis` に値 `Hypertension`
- キー `medications` にリスト `['Lisinopril']` を値として
- キー `last_visit_id` に値 `V2301`

# --hints--

辞書にはキー `diagnosis` があり、その値は文字列 `Hypertension` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("diagnosis", None) == "Hypertension"
`) })
```

辞書にはキー `medications` があり、その値はリスト `['Lisinopril']` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("medications", None) == ['Lisinopril']
`) })
```

辞書にはキー `last_visit_id` があり、その値は文字列 `V2301` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("last_visit_id", None) == "V2301"
`) })
```

辞書にはキー `patient_id` があり、その値は文字列 `P1001` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("patient_id", None) == "P1001"
`) })
```

辞書にはキー `age` があり、その値は整数 `34` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("age", None) == 34
`) })
```

辞書にはキー `gender` があり、その値は文字列 `Female` である必要があります。

```js
({ test: () => runPython(`
assert medical_records[0].get("gender", None) == "Female"
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
--fcc-editable-region--
        'gender': 'Female'
        
--fcc-editable-region--
    }    
]
```
