---
id: 6842ba0086201089ea268c92
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

前のステップで使ったのと同じ構造で、`medical_records` リストには他の患者のデータがすでに入っています。自由に確認してください。

次に、データセットを検証する関数を書き始めます。`validate` という名前の関数を作成し、パラメータを1つ `data` にしてください。

データがリストかタプルのどちらかであることを確認したいので、`validate` 関数内で `is_sequence` という変数を宣言し、`isinstance` を呼び出した結果を代入してください。`data` を最初の引数として渡し、`list` と `tuple` を含むタプルを2番目の引数として渡してください。

# --hints--

`validate` という名前の関数があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).has_function("validate")
`) })
```

`validate` 関数はパラメータを1つ `data` にしてください。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_args("data")
`) })
```

`is_sequence` 関数内に `validate` という変数があるはずです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("validate").has_variable("is_sequence")
`) })
```

`isinstance(data, (list, tuple))` 変数に `is_sequence` を代入してください。

```js
({ test: () => runPython(`
_cond = _Node(_code).find_function("validate").find_variable("is_sequence")
assert _cond.is_equivalent("is_sequence = isinstance(data, (list, tuple))") or _cond.is_equivalent("is_sequence = isinstance(data, (tuple, list))")
`) })
```

# --seed--

## --seed-contents--

```py
medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

--fcc-editable-region--

--fcc-editable-region--
```
