---
id: 684fe1be09b05ad785832546
title: ステップ 29
challengeType: 20
dashedName: step-29
---

# --description--

正規表現が文字 `p` に続く1つ以上の数字にマッチするようになったので、最後に確認すべきことは文字列に余分な文字が含まれていないことです。

これを確認するには、`re` モジュールの別の関数を使うことができます。`fullmatch` 関数は、正規表現パターンが文字列全体にマッチした場合にマッチオブジェクトを返し、そうでなければ `None` を返します。

```py
import re

book = "Fahrenheit 451"
print(re.fullmatch('\d+', book)) #None

print(re.fullmatch('Fahrenheit \d+', book))
# <re.Match object; span=(0, 14), match='Fahrenheit 451'>
```

`search` の呼び出しを、同じ引数を使って `fullmatch` の呼び出しに置き換えてください。

# --hints--

`re.search('p\d+', patient_id, re.IGNORECASE)` を `re.fullmatch('p\d+', patient_id, re.IGNORECASE)` に置き換えるべきです。

```js
({ test: () => runPython(`
assert _Node(_code).find_function("find_invalid_records").find_variable("constraints").is_equivalent("constraints = {'patient_id': isinstance(patient_id, str) and re.fullmatch('p\\d+', patient_id, re.IGNORECASE)}")
`) })
```

# --seed--

## --seed-contents--

```py
import re

medical_records = [
    {
        'patient_id': 'P1001',
        'age': 34,
        'gender': 'Female',
        'diagnosis': 'Hypertension',
        'medications': ['Lisinopril'],
        'last_visit_id': 'V2301',
    },
    {
        'patient_id': 'p1002',
        'age': 47,
        'gender': 'male',
        'diagnosis': 'Type 2 Diabetes',
        'medications': ['Metformin', 'Insulin'],
        'last_visit_id': 'v2302',
    },
    {
        'patient_id': 'P1003',
        'age': 29,
        'gender': 'female',
        'diagnosis': 'Asthma',
        'medications': ['Albuterol'],
        'last_visit_id': 'v2303',
    },
    {
        'patient_id': 'p1004',
        'age': 56,
        'gender': 'Male',
        'diagnosis': 'Chronic Back Pain',
        'medications': ['Ibuprofen', 'Physical Therapy'],
        'last_visit_id': 'V2304',
    }
]

def find_invalid_records(
    patient_id, age, gender, diagnosis, medications, last_visit_id
):
    constraints = {
--fcc-editable-region--
        'patient_id': isinstance(patient_id, str) and re.search('p\d+', patient_id, re.IGNORECASE)
--fcc-editable-region--
    }
    return constraints

def validate(data):
    is_sequence = isinstance(data, (list, tuple))

    if not is_sequence:
        print('Invalid format: expected a list or tuple.')
        return False
        
    is_invalid = False
    key_set = set(
        ['patient_id', 'age', 'gender', 'diagnosis', 'medications', 'last_visit_id']
    )

    for index, dictionary in enumerate(data):
        if not isinstance(dictionary, dict):
            print(f'Invalid format: expected a dictionary at position {index}.')
            is_invalid = True

        if set(dictionary.keys()) != key_set:
            print(
                f'Invalid format: {dictionary} at position {index} has missing and/or invalid keys.'
            )
            is_invalid = True

    if is_invalid:
        return False
    print('Valid format.')
    return True

validate(medical_records)
print(find_invalid_records(**medical_records[0]))
```
