---
id: 695774002591bbc5f8cf3e53
title: 할인 적용 함수 만들기
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

이 실습에서는 백분율 할인 적용 후 상품의 최종 가격을 계산하는 함수를 작성합니다.

예를 들어, 상품 가격이 `50`이고 `20`의 할인이 적용되면, 할인 금액은 `10`이고 최종 가격은 `40`입니다.

**목표:** 아래 사용자 스토리를 충족하여 실습을 완료하세요.

**사용자 스토리:**

1. `apply_discount`라는 이름의 함수를 정의해야 합니다.

2. `apply_discount` 함수는 정확히 두 개의 매개변수 `price`와 `discount`를 받아야 합니다.

3. `price`가 숫자(`int` 또는 `float`)가 아니면 함수는 문자열 `The price should be a number`을 반환해야 합니다.

4. `discount`가 숫자(`int` 또는 `float`)가 아니면 함수는 문자열 `The discount should be a number`를 반환해야 합니다.

5. `price`가 `0`보다 작거나 같으면 함수는 문자열 `The price should be greater than 0`을 반환해야 합니다.

6. `discount`가 `0`보다 작거나 `100`보다 크면 함수는 문자열 `The discount should be between 0 and 100`를 반환해야 합니다.

7. 두 입력이 모두 유효하면 함수는 가격의 백분율로 할인을 계산해야 합니다.

8. 함수는 할인을 적용한 후 최종 가격을 반환해야 합니다.

# --hints--

`apply_discount`라는 이름의 함수가 있어야 합니다.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

`apply_discount` 함수는 두 개의 매개변수 `price`와 `discount`를 받아야 합니다.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

`apply_discount`가 숫자(`price` 또는 `int`)가 아닌 `float` (첫 번째 인수)로 호출되면 `The price should be a number`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

`apply_discount`가 숫자(`discount` 또는 `int`)가 아닌 `float` (두 번째 인수)로 호출되면 `The discount should be a number`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

`apply_discount`가 `price`보다 작거나 같은 `0`로 호출되면 `The price should be greater than 0`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

`apply_discount`가 `discount`가 `0`보다 작거나 `100`보다 클 때 호출되면 `The discount should be between 0 and 100`를 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)`는 `80`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)`는 `100`를 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)`는 `50`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

`apply_discount`가 `100`의 할인을 적용하여 호출되면 `0`를 반환해야 합니다.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)`는 `59.6`을 반환해야 합니다.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
