---
id: 68e51d8f163d3f8b70f9c4e5
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

このワークショップでは、`includes()` と `slice()` メソッドを使う練習をします。

まず、`fccSentence` という名前の変数を作成し、文字列 `"freeCodeCamp is a great place to learn web development."` を代入してください。

# --hints--

`fccSentence` という名前の変数を宣言してください。

```js
assert.exists(fccSentence);
```

`fccSentence` 変数は文字列である必要があります。

```js
assert.isString(fccSentence);
```

`fccSentence` 変数には文字列 `"freeCodeCamp is a great place to learn web development."` を代入してください。

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
