---
id: 68e6172975008fea742a63cc
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

`includes()` メソッドは大文字と小文字を区別するため、ケースセンシティブです。

`hasLowercaseFCC` という名前の変数を作成してください。次に、`includes()` メソッドを使って `fccSentence` に `"freecodecamp"`（すべて小文字）が含まれているかどうかを確認した結果を代入してください。

その後、テンプレートリテラルか文字列の連結のどちらかを使って、メッセージ `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` をコンソールに出力してください。`<hasLowercaseFCC>` は変数の実際の値に置き換えてください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`hasLowercaseFCC` という名前の変数を宣言してください。

```js
assert.exists(hasLowercaseFCC);
```

`fccSentence.includes("freecodecamp")` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")` の結果を `hasLowercaseFCC` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

`fccSentence.includes("freecodecamp")` の結果は `false` であるべきです。

```js
assert.isFalse(hasLowercaseFCC);
```

`fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` を含むメッセージをログに出力してください。`<hasLowercaseFCC>` は変数の実際の値に置き換えてください。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
