---
id: 68e6172975008fea742a63ce
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

`slice()` メソッドは文字列の一部を抽出します。2つの引数を取ります：開始インデックス（含む）と終了インデックス（含まない）です。

例を示します：

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

文字列のインデックスは0から始まります。`message` 文字列では、インデックス11の文字は `f`（"freeCodeCamp" の始まり）です。

`platform` という名前の変数を作成してください。そして `message.slice(11, 23)` を使って `"freeCodeCamp"` を抽出した結果を代入してください。

# --hints--

`platform` という名前の変数を宣言してください。

```js
assert.exists(platform);
```

`platform` 変数は文字列であるべきです。

```js
assert.isString(platform);
```

`message.slice(11, 23)` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

`message.slice(11, 23)` の結果を `platform` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

`platform` 変数は `"freeCodeCamp"` と等しくなるべきです。

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
