---
id: 68e6172975008fea742a63ca
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

テンプレートリテラルか文字列の連結のどちらかを使って、メッセージ `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` をコンソールに出力してください。`<hasFreeCodeCamp>` は変数の実際の値に置き換えてください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

メッセージ `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` をコンソールに出力してください。`<hasFreeCodeCamp>` は変数の実際の値に置き換える必要があります。

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
