---
id: 68e6172975008fea742a63d0
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`slice()` を使って文字列の先頭から抽出できます。

`greetingWord` という名前の変数を作成してください。それから、`message.slice(0, 7)` を使って最初の単語 `"Welcome"` を抽出した結果を代入してください。

次に、テンプレートリテラルか文字列の連結のどちらかを使って、メッセージ `The first word is "<greetingWord>".` をコンソールに出力してください。`<greetingWord>` は変数の実際の値に置き換えてください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`greetingWord` という名前の変数を宣言してください。

```js
assert.exists(greetingWord);
```

`greetingWord` 変数は文字列であるべきです。

```js
assert.isString(greetingWord);
```

`message.slice(0, 7)` を使ってください。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

`message.slice(0, 7)` の結果を `greetingWord` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

`greetingWord` 変数は `"Welcome"` と等しくなるべきです。

```js
assert.strictEqual(greetingWord, "Welcome");
```

メッセージ `The first word is "<greetingWord>".` をログに出力してください。`<greetingWord>` は変数の実際の値に置き換えてください。

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
