---
id: 68e6172975008fea742a63cd
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

これから文字列の一部を抽出する`slice()`メソッドを学びます。

`message`という名前の変数を作成し、文字列`"Welcome to freeCodeCamp!"`を代入してください。

その後、`console.log()`を使って`"Here are some examples of the slice() method:"`をコンソールに出力してください。

# --hints--

`message`という名前の変数を宣言してください。

```js
assert.exists(message);
```

`message`変数は文字列である必要があります。

```js
assert.isString(message);
```

`message`変数は`"Welcome to freeCodeCamp!"`である必要があります。

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

`"Here are some examples of the slice() method:"`をコンソールに出力してください。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
