---
id: 68e6172875008fea742a63c9
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

前のレッスンで学んだように、`includes()` メソッドは文字列に特定の部分文字列が含まれているかをチェックし、`true` または `false` を返します。

`includes()` メソッドを使った例を示します。

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

`hasFreeCodeCamp` という名前の変数を作成してください。次に、`fccSentence` に対して `includes()` メソッドを使い、`"freeCodeCamp"` が含まれているかをチェックした結果を `hasFreeCodeCamp` に代入してください。

# --hints--

`hasFreeCodeCamp` という名前の変数を宣言してください。

```js
assert.exists(hasFreeCodeCamp);
```

`fccSentence.includes("freeCodeCamp")` を使う必要があります。

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")` の結果を `hasFreeCodeCamp` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

`fccSentence.includes("freeCodeCamp")` の結果は `true` になるはずです。

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
