---
id: 68e6172975008fea742a63d1
title: ステップ 11
challengeType: 1
dashedName: step-11
---

# --description--

`slice()` メソッドは負のインデックスを使って文字列の末尾から数えることができます。`-1` は最後の文字を指します。

例を示します：

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

`endPunctuation` という名前の変数を作成してください。次に、`message.slice(-1)` を使って最後の文字を抽出した結果を代入してください。

その後、テンプレートリテラルか文字列の連結を使って、メッセージ `The ending punctuation mark is a "<endPunctuation>"` をコンソールに出力してください。`<endPunctuation>` は変数の実際の値に置き換えてください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`endPunctuation` という名前の変数を宣言してください。

```js
assert.exists(endPunctuation);
```

`endPunctuation` 変数は文字列であるべきです。

```js
assert.isString(endPunctuation);
```

`message.slice(-1)` を使うべきです。

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

`message.slice(-1)` の結果を `endPunctuation` 変数に代入してください。

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

`endPunctuation` 変数は `"!"` と等しくなるべきです。

```js
assert.strictEqual(endPunctuation, "!");
```

メッセージ `The ending punctuation mark is a "<endPunctuation>"` をログに出力してください。`<endPunctuation>` は変数の実際の値に置き換えてください。

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
