---
id: 68ef6e42fa6e8a9a18a401fa
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

`console.log()` を使って `"Here are some examples of the includes() method:"` をコンソールに出力してください。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

`"Here are some examples of the includes() method:"` をコンソールに出力する必要があります。

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
