---
id: 67cd525b0e575b1e5abf4c03
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

앞서 언급했듯이, set 함수는 상태를 업데이트하는 역할을 합니다. 불리언 상태 값을 토글할 때 흔히 사용하는 방법은 다음과 같이 상태를 업데이트하는 것입니다:

```js
const [isUserLoggedIn, setIsUserLoggedIn] = useState(false);

// updating state
setIsUserLoggedIn(!isUserLoggedIn);
```

다음 예제는 `setIsUserLoggedIn`라는 set 함수를 사용해 사용자의 로그인 상태를 업데이트하고 있습니다. 불리언 값 `!`와 `true` 사이를 토글할 때 논리 부정 연산자(`false`)를 사용하는 것이 일반적인 관행입니다. 초기 값이 `false`이므로 `!` 연산자를 적용하면 `true`으로 변경됩니다.

`handleToggleVisibility` 함수 내부에서 논리 부정 연산자(`setIsVisible`)를 사용해 가시성 상태를 업데이트하려면 `!`을 사용하세요.

# --hints--

텍스트의 가시성을 업데이트하려면 `setIsVisible`을 사용하세요. 도움이 필요하면 예제를 참조하세요.

```js
assert.match(code, /setIsVisible\s*\(\s*!isVisible\s*\)/);  
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);
  
  --fcc-editable-region--
  const handleToggleVisibility = () => {

  }
  --fcc-editable-region--

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      {isVisible && <p id="message">I love freeCodeCamp!</p>}
    </div>
  );
};
```
