---
id: 67cd4cf8b5fe391b324ad11f
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

이제 `isVisible` 상태 변수를 만들었으니, 이를 사용하기 시작할 차례입니다.

이전 수업에서 특정 조건에 따라 다른 내용을 보여주는 인라인 조건부 렌더링에 대해 배웠습니다. 논리 AND 연산자(`&&`)를 사용해 다음과 같이 조건부로 텍스트를 렌더링하는 것이 일반적인 패턴입니다:

```jsx
function Notification({ message }) {
  return (
    <div>
      {message && <p>{message}</p>}
    </div>
  );
}
```

이 예제에서, 메시지가 참값일 때만 단락 요소가 표시됩니다. "참값"은 불리언 문맥에서 `true`로 평가되는 모든 값을 의미한다는 점을 기억하세요.

이번 단계에서는 단락 요소를 조건부로 렌더링해야 합니다. 올바르게 하면 단락이 페이지에 더 이상 보이지 않아야 합니다.

# --hints--

단락이 `isVisible`가 `true`일 때만 보이도록 조건부 렌더링하세요. 도움이 필요하면 예제를 다시 참고하세요.

```js
  {
    React.useState = (arg) => ([false, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.notExists(pEl);
  }
  {
    React.useState = (arg) => ([true, () => {}]);
    const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;
    const exports = {};
    const _a = eval(script);
    const div = await __helpers.prepTestComponent(exports.ToggleApp);
    const pEl = div.querySelector("p#message");
    assert.exists(pEl);
  }
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  const [isVisible, setIsVisible] = useState(false);

  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      --fcc-editable-region--
      <p id="message">I love freeCodeCamp!</p>
      --fcc-editable-region--
    </div>
  );
};
```
