---
id: 67cd49e05693111a41771f29
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

이전 수업에서 `useState` 훅과 배열 구조 분해를 사용해 상태 변수를 만드는 방법을 배웠습니다.

```js
const [user, setUser] = useState("Jessica");
```

상태 변수 이름을 짓는 관례는 `example`와 `setExample` 형식을 사용하는 것입니다.

`useState`는 두 개의 값을 가진 배열을 반환합니다. 첫 번째 값은 현재 상태이고, 두 번째 값은 상태를 업데이트하는 데 사용할 설정 함수입니다.

새 상태 변수와 설정 함수 이름을 각각 `isVisible`와 `setIsVisible`로 만드세요. 예제에 나온 배열 구조 분해 구문을 사용해야 합니다.

그다음 `useState` 훅에 초기값 `false`을 할당하세요.

# --hints--

`isVisible`와 `setIsVisible` 변수에는 배열 구조 분해 구문을 사용해야 합니다. 예제를 다시 참고하세요.

```js
assert.match(code, /(const|let|var)\s+\[\s*(isVisible)\s*,\s*(setIsVisible)\s*\]\s*/);
```

`useState` 훅을 사용해야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.isAtLeast(abuseState.calls.length, 1);
```

`useState` 훅에는 초기값 `false`이 있어야 합니다.

```js
  const abuseState = __helpers.spyOn(React, "useState");

  const script = [...document.querySelectorAll("script")].find((s) => s.dataset.src ===  "index.jsx").innerText;

  const exports = {};
  const _a = eval(script);
  const _b = await __helpers.prepTestComponent(exports.ToggleApp);
  assert.equal(abuseState.calls[0]?.[0], false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
<head>
  <meta charset="UTF-8" />
  <title>Toggle Visibility</title>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react/18.3.1/umd/react.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/react-dom/18.3.1/umd/react-dom.development.min.js"></script>
  <script src="https://cdnjs.cloudflare.com/ajax/libs/babel-standalone/7.26.5/babel.min.js"></script>
  <script 
    data-plugins="transform-modules-umd"
    type="text/babel"
    src="index.jsx"
  ></script>
  <link rel="stylesheet" href="styles.css" />
</head>
<body>
  <div id="root"></div>
  <script
    data-plugins="transform-modules-umd"
    type="text/babel"
    data-presets="react"
    data-type="module"
  >
    import { ToggleApp } from "./index.jsx";
    ReactDOM.createRoot(document.getElementById("root")).render(<ToggleApp />);
  </script>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f4f4f4;
}

#toggle-container {
  text-align: center;
  background: white;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 0 10px rgba(0, 0, 0, 0.1);
}

#toggle-button {
  padding: 10px 20px;
  border: none;
  background: #007BFF;
  color: white;
  border-radius: 5px;
  cursor: pointer;
  transition: background 0.3s ease;
}

#toggle-button:hover {
  background: #0056b3;
}

#message {
  margin-top: 20px;
  font-size: 1.125rem;
  color: #333;
}
```

```jsx
const { useState } = React;

export const ToggleApp = () => {
  --fcc-editable-region--

  --fcc-editable-region--
  
  return (
    <div id="toggle-container">
      <button id="toggle-button">Message</button>
      <p id="message">I love freeCodeCamp!</p>
    </div>
  );
};
```
