---
id: 5f33294a6af5e9188dbdb8f3
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Ni wakati wa kuongeza baadhi ya maudhui ya menyu. Ongeza kipengele cha `main` ndani ya kipengele kilichopo cha `body`. Hatimaye kitakuwa na taarifa za bei kuhusu kahawa na vitafunwa vinavyotolewa na mkahawa.

# --hints--

Msimbo wako unapaswa kuwa na lebo ya kufungua ya `<main>`.

```js
assert(code.match(/<main>/i));
```

Msimbo wako unapaswa kuwa na lebo ya kufunga ya `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Haupasi kubadilisha kipengele chako cha `body`. Hakikisha hujafuta lebo yako ya kufunga kwa bahati mbaya.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Leboa ya `main` inapaswa kuwa ndani ya lebo ya `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
