---
id: 5f33071498eb2472b87ddee4
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika mradi huu, utafundishwa misingi ya CSS (Cascading Style Sheets) kwa kujenga menyu ya kafé. CSS ni lugha inayotumika kupamba hati ya HTML. Inafafanua jinsi vipengele vya HTML vinavyopaswa kuonyeshwa kwenye skrini.

Kama ulivyojifunza katika hatua chache za mwisho za Cat Photo App, kuna muundo wa msingi unaohitajika kuanza kujenga ukurasa wako wa mtandao. Kila hati ya HTML inapaswa kuwa na tamko la `DOCTYPE` na kipengele cha `html`. `DOCTYPE` inaeleza kivinjari toleo gani la HTML hati hiyo iko. Na kipengele cha `html` kinakilisha kipengele cha mzizi kinachojumuisha vipengele vyote vingine.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Ongeza lebo ya `<!DOCTYPE html>`, na kipengele cha `html` chenye sifa ya kipengele ya `lang` yenye thamani ya `en`.

# --hints--

Unapaswa kuwa na tamko la `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Unapaswa kuwa na lebo ya ufunguzi ya `<html lang="en">` yenye lugha imewekwa kuwa Kiingereza.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Unapaswa kuwa na lebo ya kufunga ya `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
