---
id: 5f3326b143638ee1a09ff1e3
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Ili kujiandaa kuunda maudhui halisi, ongeza kipengele cha `body` chini ya kipengele cha `head`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<body>`.

```js
assert(code.match(/<body>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `</body>`.

```js
assert(code.match(/<\/body>/i));
```

Hupaswi kubadilisha kipengele chako cha `head`. Hakikisha hujafuta lebo yako ya kufunga.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Kipengele chako cha `body` kinapaswa kuja baada ya kipengele chako cha `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
