---
id: 5f3313e74582ad9d063e3a38
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ongeza kipengele cha `head` ndani ya kipengele cha `html`, ili uweze kuongeza kipengele cha `title`. Maandishi ya kipengele cha `title` yanapaswa kuwa `Cafe Menu`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `<head>`.

```js
assert(code.match(/<head>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `</head>`.

```js
assert(code.match(/<head>/i));
```

Unapaswa kuwa na lebo ya ufunguzi ya `<title>`.

```js
assert(code.match(/<title>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `</title>`.

```js
assert(code.match(/<\/title>/i));
```

Kipengele chako cha `title` kinapaswa kupangwa ndani ya kipengele chako cha `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Kipengele chako cha `title` kinapaswa kuwa na maandishi ya `Cafe Menu`. Huenda unahitaji kukagua tahajia yako.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
