---
id: 5f331e55dfab7a896e53c3a1
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

`title` ni mojawapo ya vipengele kadhaa vinavyotoa taarifa za ziada ambazo hazionekani kwenye ukurasa wa mtandao, lakini ni muhimu kwa mashine za utafutaji au jinsi ukurasa unavyoonyeshwa.

Ndani ya kipengele cha `head`, panga ndani kipengele cha `meta` chenye sifa ya kipengele inayoitwa `charset` iliyowekwa kwa thamani ya `utf-8` ili kumwambia kivinjari jinsi ya kusimbua herufi za ukurasa.

Kumbuka kuwa vipengele vya `meta` ni vipengele visivyo na mwisho vinavyohitaji lebo ya kumalizia.

# --hints--

Unapaswa kuwa na lebo ya `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Kipengele cha `meta` ni kipengele kisicho na mwisho, hakipaswi kuwa na lebo ya kumalizia `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Leboa yako ya `meta` inapaswa kuwa na sifa ya kipengele ya `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Sifa yako ya kipengele ya `charset` inapaswa kuwa na thamani ya `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
