---
id: 5f332a88dc25a0fd25c7687a
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Jina la mkahawa ni `CAMPER CAFE`. Ongeza kipengele cha `h1` ndani ya kipengele chako cha `main`. Mpe jina la mkahawa kwa herufi kubwa ili liweze kuonekana vizuri.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `<h1>`.

```js
assert(code.match(/<h1>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `</h1>`.

```js
assert(code.match(/<\/h1>/i));
```

Hupaswi kubadilisha kipengele chako cha `main` kilichopo.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Kipengele chako cha `h1` kinapaswa kupangwa ndani ya kipengele chako cha `main`.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `CAMPER CAFE`.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```
