---
id: 5f3477ae34c1239cafe128be
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Sasa una vichaguzi vya aina vitatu vyenye mtindo uleule. Unaweza kuongeza kundi lile lile la mitindo kwa vipengele vingi kwa kuunda orodha ya vichaguzi. Kila kichaguzi kimegawanywa kwa koma kama hii:

```css
selector1, selector2 {
  property: value;
}
```

Futa vichaguzi vya aina vitatu vilivyopo na ubadilishe na orodha moja ya vichaguzi inayoweka maandishi katikati kwa vipengele vya `h1`, `h2`, na `p`.

# --hints--

Unapaswa kutumia kichaguzi cha aina kimoja kwa vipengele vyote vitatu, `h1, h2, p`. Hakikisha unafuata mpangilio huo.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Unapaswa kuwa na kichaguzi kimoja tu katika kipengele chako cha `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
