---
id: 5f3477aefa51bfc29327200b
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Umeipamba vipengele vitatu kwa kuandika CSS ndani ya lebo za `style`. Hii inafanya kazi, lakini kwa kuwa kutakuwa na mitindo mingi zaidi, ni bora kuweka mitindo yote katika faili tofauti na kuunganisha nayo.

Faili tofauti ya `styles.css` imeundwa kwa ajili yako. Unaweza kubadilisha kati ya mafaili kwa kutumia tabo zilizo juu ya mhariri.

Anza kwa kuandika upya mitindo uliyotengeneza katika faili ya `styles.css`. Hakikisha haujumuishi lebo za kufungua na kufunga za `style`.

# --hints--

Faili yako ya `styles.css` inapaswa kuwa na kichaguzi cha aina cha `h1, h2, p`.

```js
assert(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, "").match(__helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/)));
```
 
Kichaguzi chako kinapaswa kuweka sifa ya `text-align` kuwa `center`.

```js
assert(code.slice(code.indexOf('</html>')).match(/text-align:\s*center;?/));
```

Msimbo wako haupaswi kuwa na lebo za `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
