---
id: 5f3477ae9675db8bb7655b30
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Katika hatua iliyopita, ulitumia <dfn>kichaguzi cha aina</dfn> kuunda mtindo wa kipengele cha `h1`. Weka maudhui ya vipengele vya `h2` na `p` katikati kwa kuongeza kichaguzi kipya cha aina kwa kila kimoja kwenye kipengele kilichopo cha `style`.

# --hints--

Haupaswi kubadilisha kichaguzi kilichopo cha `h1`.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

Haupaswi kuongeza lebo mpya ya `style`. Ongeza sheria mpya za CSS kwenye lebo iliyopo ya `style`.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

Unapaswa kuongeza kichaguzi kipya cha `h2`.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

Unapaswa kuongeza kichaguzi kipya cha `p`.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

Kipengele chako cha `h1` kinapaswa kuwa na `text-align` ya `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

Kipengele chako cha `h2` kinapaswa kuwa na `text-align` ya `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

Kipengele chako cha `p` kinapaswa kuwa na `text-align` ya `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
