---
id: 5f3477cb2e27333b1ab2b955
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--

Sasa unahitaji kuunganisha faili la `styles.css`, ili mitindo itumike tena. Ndani ya kipengele cha `head`, ongeza kipengele cha `link`. Mpe sifa ya kipengele ya `rel` yenye thamani ya `"stylesheet"` na sifa ya kipengele ya `href` yenye thamani ya `"styles.css"`.

Kumbuka kwamba kipengele cha `link` ni kipengele kisicho na maudhui.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
const link = document.querySelector('link');
assert.isNotNull(link);
```

Kipengele cha `link` ni kipengele kisicho na maudhui, hakipaswi kuwa na lebo ya kufunga `</link>`.

```js
assert.notMatch(code, /<\/link>/i);
```

Hupaswi kubadilisha kipengele chako cha `head` kilichopo. Hakikisha hujafuta lebo ya kufunga.

```js
const headElementCount = document.querySelectorAll('head')?.length;
assert.strictEqual(headElementCount, 1);
```

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
const linkElementCount = document.querySelectorAll('link')?.length;
assert.strictEqual(linkElementCount, 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const link = document.querySelector('head > link');
assert.isNotNull(link);
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
const linkRelValue = document.querySelector('link')?.getAttribute('rel');
assert.strictEqual(linkRelValue, 'stylesheet');
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
const linkHrefValue = document.querySelector('link')?.dataset?.href;
assert.strictEqual(linkHrefValue, 'styles.css');

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
