---
id: 5f3477ae8466a9a3d2cc953c
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Sasa baada ya kuwa na CSS katika faili la `styles.css`, endelea na ondoa kipengele cha `style` pamoja na maudhui yake yote. Mara kitakapoondolewa, maandishi yaliyo katikati yatahamia upande wa kushoto tena.

# --hints--

Haupaswi kuwa na lebo za `style` katika msimbo wako.

```js
assert(!code.match(/style/i));
```

Haupaswi kuwa na vichaguzi vya CSS katika faili lako la HTML.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
