---
id: 5f344fbc22624a2976425065
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Unda kipengele cha `h2` ndani ya kipengele cha `section` na kimpatie maandishi ya `Coffee`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Unapaswa kuwa na lebo ya kufunga ya `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Hupaswi kubadilisha kipengele chako cha `section` kilichopo. Hakikisha hujafuta lebo ya kufunga.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Kipengele chako cha `h2` kinapaswa kuwa ndani ya kipengele chako cha `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Kipengele chako cha `h2` kinapaswa kuwa na maandishi ya `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
