---
id: 5f356ed656a336993abd9f7c
title: Hatua 25
challengeType: 0
dashedName: step-25
---

# --description--

Ifuatayo, unataka kuweka `#menu` katikati kwa usawa. Unaweza kufanya hivi kwa kuweka sifa zake za `margin-left` na `margin-right` kuwa `auto`. Fikiria ukingo kama nafasi isiyoonekana inayozunguka kipengele. Kwa kutumia sifa hizi mbili za ukingo, weka kipengele cha `#menu` katikati ndani ya kipengele cha `body`.

# --hints--

Unapaswa kuweka sifa ya `margin-left` kuwa `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-left'] === 'auto');
assert(hasMargin);
```

Unapaswa kuweka sifa ya `margin-right` kuwa `auto`.

```js
const hasMargin = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['margin-right'] === 'auto');
assert(hasMargin);
```

Unapaswa kuweka sifa za `margin-left` na `margin-right` za `#menu` yako kuwa `auto`.

```js
const divMarginRight = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-right');
const divMarginLeft = new __helpers.CSSHelp(document).getStyle('#menu')?.getPropertyValue('margin-left');
assert(divMarginRight === 'auto');
assert(divMarginLeft === 'auto');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
}
--fcc-editable-region--
```
