---
id: 5f35e5c44359872a137bd98f
title: Hatua 28
challengeType: 0
dashedName: step-28
---

# --description--

Kwa kuwa bidhaa kuu ya kahawa inayouzwa ni kahawa, unaweza kutumia picha ya maharagwe ya kahawa kama mandhari ya ukurasa.

Ondoa maoni na yaliyomo ndani ya kichaguzi cha aina cha `body`. Baada ya hapo, ongeza sifa ya `background-image` na uweke thamani yake kuwa `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Unapaswa kuondoa sifa ya `background-color` iliyotangazwa kama maoni.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

Kichaguzi chako cha `body` hakipaswi kuwa na maoni yoyote.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

Unapaswa kuweka sifa ya `background-image` kuwa `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

Kipengele chako cha `body` kinapaswa kuwa na picha ya mandhari ya maharagwe ya kahawa.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
