---
id: 5f356ed63c7807a4f1e6d054
title: Hatua 21
challengeType: 0
dashedName: step-21
---

# --description--

Lengo sasa ni kufanya `div` isichukue upana mzima wa ukurasa. Sifa ya CSS `width` ni kamili kwa hili.

Unaweza kutumia kichaguzi cha `id` kulenga kipengele maalum chenye sifa ya `id`. Kichaguzi cha `id` kimefafanuliwa kwa kuweka alama ya hash `#` moja kwa moja mbele ya thamani ya `id` ya kipengele. Kwa mfano, kama kipengele kina `id` ya `cat` basi ungetaka kulenga kipengele hicho hivi:

```css 
#cat {
  width: 250px;
}
```

Tumia kichaguzi cha `#menu` kumpa kipengele chako upana wa `300px`.

# --hints--

Unapaswa kuwa na kichaguzi cha `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

Unapaswa kuweka sifa ya `width` kuwa `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

`div` yako inapaswa kuwa na upana wa 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```
