---
id: 5f356ed60785e1f3e9850b6e
title: Hatua 24
challengeType: 0
dashedName: step-24
---

# --description--

Sasa ni rahisi kuona kuwa maandishi yamewekwa katikati ndani ya kipengele cha `#menu`. Kwa sasa, upana wa kipengele cha `#menu` umeelezwa kwa pikseli (`px`).

Badilisha thamani ya sifa ya `width` kuwa `80%`, ili iwe asilimia 80 ya upana wa kipengele chake cha mzazi (`body`).

# --hints--

Unapaswa kuweka sifa ya `width` kuwa `80%`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

Haupasi kuwa na sifa ya `width` yenye thamani ya `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```
