---
id: 66d75dd0aa65a71600dc669b
title: 建置一個庫存管理程式
challengeType: 26
dashedName: build-an-inventory-management-program
---

# --description--

在這個實驗中，你將建置一個庫存管理程式，允許你為庫存添加、更新、查找和移除產品。你將使用一個物件陣列來表述你的庫存，其中每個物件將以 name 和 quantity 作為鍵。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該宣告一個名為 `inventory` 的空陣列，用來儲存產品物件，該物件具有一個鍵 `name`，其值為小寫字串，以及一個鍵 `quantity`，其值為整數。
1. 你應該創建一個名為 `findProductIndex` 的函式（程式），該函式以產品名稱作為引數，並傳回對應產品物件在 `inventory` 陣列中的索引。該函式應始終使用產品名稱的小寫形式來執行搜尋。如果找不到產品，該函式應傳回 `-1`。
1. 你應該創建一個名為 `addProduct` 的函式（程式），該函式（程式）以產品物件作為其引數。
1. 如果產品已存在於庫存中，`addProduct` 函式（程式）應該透過將傳入函式的數量加到目前的 `quantity` 來更新其 `quantity` 值，並在主控台記錄產品名稱，後面接一個空格和 `quantity updated`。
1. 如果產品不在庫存中，`addProduct` 函式（程式）應該將產品推入 `inventory` 陣列，並將產品名稱記錄到主控台，後面接一個空格和 `added to inventory`。
1. 你應該創建一個名為 `removeProduct` 的函式（程式），該函式以產品名稱和數量作為其引數。
1. `removeProduct` 函式（程式）應該從庫存中對應的產品物件中減去傳入的數量，並將字串 `Remaining <product-name> pieces: <product-quantity>` 記錄到主控台，其中 `<product-name>` 應該被產品名稱取代，`<product-quantity>` 應該被產品數量取代。
1. 如果減法後的數量為零，`removeProduct` 應該從庫存中移除該產品物件。如果庫存中的數量不足以執行減法，`removeProduct` 函式（程式）應該將字串 `Not enough <product-name> available, remaining pieces: <product-quantity>` 記錄到主控台。
1. 如果要移除的產品不在庫存中，`removeProduct` 函式（程式）應該在主控台記錄 `<product-name> not found`。

**注意：** 為了避免衝突，在執行測試時只保留使用者故事中提到的日誌記錄。

# --hints--

你應該宣告一個名為 `inventory` 的空陣列。

```js
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+inventory\s*=\s*\[\s*\]\s*;?/)
```

你應該有一個名為 `findProductIndex` 的函式（程式）。

```js
assert.isFunction(findProductIndex)
```

`findProductIndex("flour")` 應該傳回 `inventory` 陣列中 `name` 與 `"flour"` 相等的物件的索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("flour"), 0);
assert.strictEqual(findProductIndex("rice"), 1);
assert.strictEqual(findProductIndex("sugar"), 2);
```

`findProductIndex("FloUr")` 應該傳回 `inventory` 陣列中 `name` 與 `"flour"` 相等的物件索引。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("FLour"), 0);
assert.strictEqual(findProductIndex("RICE"), 1);
assert.strictEqual(findProductIndex("suGar"), 2);
```

當在 `inventory` 陣列中找不到 `name` 與 `"flour"` 相等的物件時，`findProductIndex("Flour")` 應該傳回 `-1`。

```js
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
assert.strictEqual(findProductIndex("salt"), -1)
```

你應該有一個名為 `addProduct` 的函式（程式）。

```js
assert.isFunction(addProduct)
```

`addProduct({name: "FLOUR", quantity: 5})` 應該將 `5` 添加到 `inventory` 陣列中 `name` 與 `"flour"` 相等的物件的 `quantity` 值。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'Flour', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'SuGar', quantity: 5});
const expected = JSON.stringify([{name: 'flour', quantity: 20}, {name: 'rice', quantity: 15}, {name: 'sugar', quantity: 10}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

當在 `inventory` 陣列中找到具有 `name` 相等於 `"flour"` 的物件時，`addProduct({name: "FLOUR", quantity: 5})` 應該記錄 `flour quantity updated`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'flOur', quantity: 5});
addProduct({name: 'RICE', quantity: 5});
addProduct({name: 'sugar', quantity: 5});
console.log = temp;
assert.lengthOf(logArr, 3);
assert.strictEqual(logArr[0], "flour quantity updated")
assert.strictEqual(logArr[1], "rice quantity updated")
assert.strictEqual(logArr[2], "sugar quantity updated")
```

當在 `inventory` 陣列中找不到 `name` 與 `"flour"` 相等的物件時，`addProduct({name: "FLOUR", quantity: 5})` 應該將 `{name: "flour", quantity: 5}` 推入 `inventory` 陣列。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SALT', quantity: 8});
addProduct({name: 'Spam', quantity: 1});
const expected = JSON.stringify([{name: 'flour', quantity: 15}, {name: 'rice', quantity: 10}, {name: 'sugar', quantity: 5}, {name: 'salt', quantity: 8}, {name: 'spam', quantity: 1}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected);
```

當庫存中找不到 `name` 相等於 `"flour"` 的物件時，`addProduct({name: "FLOUR", quantity: 5})` 應該記錄 `flour added to inventory`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
addProduct({name: 'SAlt', quantity: 5});
addProduct({name: 'SPAM', quantity: 1});
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt added to inventory");
assert.equal(logArr[1], "spam added to inventory");
```

你應該創建一個名為 `removeProduct` 的函式（程式）。

```js
assert.isFunction(removeProduct);
```

當在 `inventory` 陣列中找不到 `name` 與 `"flour"` 相等的物件時，`removeProduct("FLOUR", 5)` 應該記錄 `flour not found`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("SALT", 2);
removeProduct("Spam", 1);
console.log = temp;
assert.lengthOf(logArr, 2);
assert.equal(logArr[0], "salt not found");
assert.equal(logArr[1], "spam not found");
```

`removeProduct("FLOUR", 5)` 應該從 `inventory` 陣列中 `name` 相等於 `"flour"` 的物件的 `quantity` 值中減去 `5`。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 10);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
const expected = JSON.stringify([{name: 'flour', quantity: 5}, {name: 'rice', quantity: 5}, {name: 'sugar', quantity: 3}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 5)` 應該在 `inventory` 相等於 `[{name: "flour", quantity: 20}, {name: "rice", quantity: 5}]` 時，將 `Remaining flour pieces: 15` 記錄到主控台。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("Flour", 5);
removeProduct("RICE", 5);
removeProduct("Sugar", 2);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Remaining flour pieces: 10");
assert.equal(logArr[1], "Remaining rice pieces: 5");
assert.equal(logArr[2], "Remaining sugar pieces: 3");
```

如果減法後的 `quantity` 為零，`removeProduct` 應該從庫存中移除該產品物件。

```js
const temp = console.log;
console.log = () => {};
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 15);
removeProduct("rice", 10);
const expected = JSON.stringify([{name: 'sugar', quantity: 5}]);
console.log = temp;
assert.equal(JSON.stringify(inventory), expected)
```

`removeProduct("FLOUR", 10)` 應該在 `inventory` 與 `[{name: "flour", quantity: 5}, {name: "rice", quantity: 5}]` 相等時，記錄 `Not enough flour available, remaining pieces: 5`。

```js
const logArr = [];
const temp = console.log;
console.log = function(...args) {
    for (const arg of args) {
        logArr.push(arg);
    }
}
inventory.length = 0;
inventory.push({name: 'flour', quantity: 15});
inventory.push({name: 'rice', quantity: 10});
inventory.push({name: 'sugar', quantity: 5});
removeProduct("flour", 20);
removeProduct("RICE", 20);
removeProduct("Sugar", 20);
console.log = temp;
assert.lengthOf(logArr, 3);
assert.equal(logArr[0], "Not enough flour available, remaining pieces: 15")
assert.equal(logArr[1], "Not enough rice available, remaining pieces: 10")
assert.equal(logArr[2], "Not enough sugar available, remaining pieces: 5")
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const inventory = [];

const findProductIndex = (productName) => {
    let productIndex = -1;
    inventory.forEach((element, index) => {
        if (element.name === productName.toLowerCase()) {
            productIndex = index;
        }
    })
    return productIndex;
}

const addProduct = (product) => {
    product.name = product.name.toLowerCase();
    const productIndex = findProductIndex(product.name);
    if (productIndex < 0) {
        inventory.push(product);
        console.log(`${product.name} added to inventory`)
        return;
    }
    inventory[productIndex].quantity += product.quantity;
    console.log(`${product.name} quantity updated`)
}

const removeProduct = (productName, productQuantity) => {
    productName = productName.toLowerCase();
    const productIndex = findProductIndex(productName);
    if (productIndex < 0) {
        console.log(`${productName} not found`);
        return;
    }
    const foundProduct = inventory[productIndex];
    if (foundProduct.quantity >= productQuantity) {
        foundProduct.quantity -= productQuantity;
        console.log(`Remaining ${productName} pieces: ${foundProduct.quantity}`)
        if (!foundProduct.quantity) {
            inventory.splice(productIndex, 1);
        }
        return;
    }
    console.log(`Not enough ${productName} available, remaining pieces: ${foundProduct.quantity}`);
}

// Example usage

const product1 = {
    name: 'flour',
    quantity: 15
}
const product2 = {
    name: 'flour',
    quantity: 10
}
const product3 = {
    name: 'sugar',
    quantity: 20
}
const product4 = {
    name: 'salt',
    quantity: 12
}

addProduct(product1);
addProduct(product2);
addProduct(product3);
addProduct(product4);
removeProduct('flour', 30);
removeProduct('flour', 20);
removeProduct('flour', 5);
```
