---
id: a0b5010f579e69b815e7c5d6
title: Створіть розумну функцію заміни слова
challengeType: 26
dashedName: build-a-smart-word-replacement-function
---

# --description--

У цій лабораторній роботі ви створите функцію, яка виконує пошук і заміну у реченні, використовуючи надані аргументи, а потім повертає нове речення.

**Примітка:** Зберігайте регістр першого символу в оригінальному слові під час заміни. Наприклад, якщо потрібно замінити слово `Book` на слово `dog`, воно має бути замінене як `Dog`

**Мета**: Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні створити функцію з назвою `myReplace`.
1. Функція `myReplace` має приймати три аргументи: рядок, слово, яке потрібно замінити, і слово, яким його потрібно замінити.
1. Функція `myReplace` має повертати новий рядок із заміною.
1. Ви повинні зберігати регістр першого символу в оригінальному слові під час заміни.

# --hints--

Ви повинні створити функцію з назвою `myReplace`.

```js
assert.isFunction(myReplace);
```

`myReplace` має приймати три аргументи.

```js
assert.lengthOf(myReplace, 3);
```

`myReplace("Let us go to the store", "store", "mall")` має повертати рядок `Let us go to the mall`.

```js
assert.deepEqual(
  myReplace('Let us go to the store', 'store', 'mall'),
  'Let us go to the mall'
);
```

`myReplace("He is Sleeping on the couch", "Sleeping", "sitting")` має повертати рядок `He is Sitting on the couch`.

```js
assert.deepEqual(
  myReplace('He is Sleeping on the couch', 'Sleeping', 'sitting'),
  'He is Sitting on the couch'
);
```

`myReplace("I think we should look up there", "up", "Down")` має повертати рядок `I think we should look down there`.

```js
assert.deepEqual(
  myReplace('I think we should look up there', 'up', 'Down'),
  'I think we should look down there'
);
```

`myReplace("This has a spellngi error", "spellngi", "spelling")` має повертати рядок `This has a spelling error`.

```js
assert.deepEqual(
  myReplace('This has a spellngi error', 'spellngi', 'spelling'),
  'This has a spelling error'
);
```

`myReplace("His name is Tom", "Tom", "john")` має повертати рядок `His name is John`.

```js
assert.deepEqual(
  myReplace('His name is Tom', 'Tom', 'john'),
  'His name is John'
);
```

`myReplace("Let us get back to more Coding", "Coding", "algorithms")` має повертати рядок `Let us get back to more Algorithms`.

```js
assert.deepEqual(
  myReplace('Let us get back to more Coding', 'Coding', 'algorithms'),
  'Let us get back to more Algorithms'
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function myReplace(str, before, after) {
  if (before.charAt(0) === before.charAt(0).toUpperCase()) {
    after = after.charAt(0).toUpperCase() + after.substring(1);
  } else {
    after = after.charAt(0).toLowerCase() + after.substring(1);
  }
  return str.replace(before, after);
}
```
