---
id: 68b94f23f7b3ee14078c89c5
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo workshop, userai l'elemento `iframe` per mostrare un video. Il codice boilerplate HTML di base è già pronto per te.

Inizia creando un elemento `h1` con il testo `iframe Video Display`.

# --hints--

Dovresti avere un elemento `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Il tuo elemento `h1` dovrebbe avere il testo `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
