---
id: 68ba9c4f1688914d72876095
title: Passaggio 4
challengeType: 0
dashedName: step-4
---

# --description--

L'elemento `iframe` prende anche un attributo `src` con un valore che indica l'URL o il percorso della risorsa da visualizzare.

Aggiungi un attributo `src` di `https://www.youtube.com/embed/I0_951_MPE0` al tuo elemento `iframe`.

A questo punto dovresti vedere il video visualizzato sulla pagina, ma ci sono altri attributi che devi aggiungere.

# --hints--

Il tuo elemento `iframe` dovrebbe avere un attributo `src` impostato su `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
