---
id: 68ba9c4f1688914d72876094
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Nella prima lezione sull'elemento `iframe` hai imparato che è un elemento sostituito proprio come `img`. Questo significa che può anche avere le proprietà `width` e `height` per determinare quanto deve essere alto e largo.

Dai al tuo elemento `iframe` un attributo `width` impostato su `560` e un attributo `height` impostato su `315`.

# --hints--

Il tuo elemento `iframe` dovrebbe avere un attributo `width` impostato su `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Il tuo elemento `iframe` dovrebbe avere un attributo `height` impostato su `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
