---
id: 68ba9c4f1688914d72876096
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Uno degli attributi Ă¨ `allow`. Ăˆ come una lista di permessi che dice al browser quali funzionalitĂ  `iframe` puĂ² usare.

Ecco un elemento `iframe` con l'attributo `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Aggiungi l'attributo `allow` con i valori `accelerometer`, `autoplay` e `clipboard-write`.

`accelerometer` permette a `iframe` di usare i sensori di movimento per rilevare cose come l'inclinazione e la rotazione del dispositivo. `autoplay` permette al video di partire automaticamente, e `clipboard-write` permette all'iframe di scrivere dati negli appunti dell'utente.

# --hints--

Il tuo elemento `iframe` dovrebbe avere un attributo `allow` impostato su `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

L'attributo `allow` del tuo elemento `iframe` dovrebbe avere `autoplay` come uno dei suoi valori.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

L'attributo `allow` del tuo elemento `iframe` dovrebbe avere `clipboard-write` come uno dei suoi valori.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
