---
id: 68ba9c4f1688914d72876098
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Il prossimo attributo che aggiungerai è `referrerpolicy`. È la regola che determina quanta dettagliata informazione condividi quando la tua pagina si collega a un'altra pagina.

Aggiungi l'attributo `referrerpolicy` e impostalo su `strict-origin-when-cross-origin`. Questo condivide l'indirizzo completo sullo stesso sito, solo il nome del sito su altri siti e niente su siti non sicuri.

# --hints--

Il tuo elemento `iframe` dovrebbe avere un attributo `referrerpolicy` impostato su `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
