---
id: 68ba9c4f1688914d72876097
title: Passaggio 6
challengeType: 0
dashedName: step-6
---

# --description--

Aggiungi `encrypted-media`, `gyroscope` e `web-share` ai valori esistenti nell'attributo `allow`.

Questi tre permetteranno di usare le estensioni per media criptati per proteggere il video, concedere l'accesso ai sensori di movimento e orientamento del dispositivo e consentire la condivisione del contenuto dell'iframe tramite i dialoghi di condivisione nativi del dispositivo.

# --hints--

L'attributo `allow` del tuo elemento `iframe` dovrebbe avere `encrypted-media` come uno dei suoi valori.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

L'attributo `allow` del tuo elemento `iframe` dovrebbe avere `gyroscope` come uno dei suoi valori.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

L'attributo `allow` del tuo elemento `iframe` dovrebbe avere `web-share` come uno dei suoi valori.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
