---
id: 66e028680eca7d21db7e1aee
title: Побудова пісочниці для регулярних виразів
challengeType: 25
dashedName: lab-regex-sandbox
demoType: onClick
---

# --description--

Для цієї лабораторної роботи у вас вже є готові CSS та HTML. Ви використаєте JavaScript, щоб увімкнути пісочницю для регулярних виразів, яка тестуватиме регулярний вираз на рядку та підсвічуватиме результати.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні отримати доступ до елементів `#pattern`, `#test-string`, `#test-btn` та `#result` і зберегти їх у змінних `regexPattern`, `stringToTest`, `testButton` та `testResult` відповідно.
1. Ви повинні отримати доступ до прапорців `#i` та `#g` і зберегти їх у змінних `caseInsensitiveFlag` та `globalFlag` відповідно.
1. У вас має бути функція з назвою `getFlags`, яка повертає рядок, що містить прапорці з увімкнених прапорців, або порожній рядок, якщо жоден із них не увімкнено.
1. Коли ви натискаєте кнопку `testButton`, якщо шаблон регулярного виразу співпадає з тестовим рядком, знайдений текст має бути підсвічений. Щоб підсвітити збіги, замініть кожен збіг на той самий текст, оточений елементом `span` з класом `highlight`. Зверніть увагу, що CSS для цього вже додано.
1. Якщо є збіги, знайдений текст має відображатися всередині `#result`. У разі кількох збігів кожен знайдений текст має бути відокремлений від наступного комою та пробілом.
1. Якщо збігів немає, у `no match` має відображатися текст `#result`.

# --hints--

Ви повинні отримати доступ до елемента `#pattern` і зберегти його у змінній `regexPattern`.

```js
assert.equal(regexPattern, document.querySelector('#pattern'))
```

Ви повинні отримати доступ до елемента `#test-string` і зберегти його у змінній `stringToTest`.

```js
assert.equal(stringToTest, document.querySelector('#test-string'))
```

Ви повинні отримати доступ до елемента `#test-btn` і зберегти його у змінній `testButton`.

```js
assert.equal(testButton, document.querySelector('#test-btn'))
```

Ви повинні отримати доступ до елемента `#result` і зберегти його у змінній `testResult`.

```js
assert.equal(testResult, document.querySelector('#result'))
```

Ви повинні отримати доступ до прапорця `#i` і зберегти його у `caseInsensitiveFlag`.

```js
assert.equal(caseInsensitiveFlag, document.querySelector('#i'))
```

Ви повинні отримати доступ до прапорця `#g` і зберегти його у змінній `globalFlag`.

```js
assert.equal(globalFlag, document.querySelector('#g'))
```

У вас має бути функція `getFlags`.

```js
assert.isFunction(getFlags)
```

Функція `getFlags` має повертати `"i"`, якщо прапорець `#i` увімкнено.

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = false;
assert.equal(getFlags(), "i")
```

Функція `getFlags` має повертати `"g"`, якщо прапорець `#g` увімкнено.

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
assert.equal(getFlags(), "g")
```

Функція `getFlags` має повертати або `"ig"`, або `"gi"`, якщо обидва прапорці увімкнено.

```js
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const flags = getFlags();
assert.lengthOf(flags, 2);
assert.include(flags, "i");
assert.include(flags, "g");
```

Функція `getFlags` має повертати порожній рядок, якщо обидва прапорці не увімкнено.

```js
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
assert.equal(getFlags(), "")
```

Коли внутрішній HTML `stringToTest` дорівнює `Gu1n34 P1g5`, значення `regexPattern` дорівнює `\d+`, і жоден прапорець не увімкнено, `stringToTest.innerHTML` має стати `Gu<span class="highlight">1</span>n34 P1g5` при натисканні кнопки `testButton`.

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n34 P1g5';
assert.equal(stringToTest.innerHTML, expected);
```

Коли внутрішній HTML `stringToTest` дорівнює `Gu1n34 P1g5`, значення `regexPattern` дорівнює `\d+`, і прапорець глобального режиму увімкнено, `stringToTest.innerHTML` має стати `Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>` при натисканні кнопки `testButton`.

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'Gu<span class="highlight">1</span>n<span class="highlight">34</span> P<span class="highlight">1</span>g<span class="highlight">5</span>';
assert.equal(stringToTest.innerHTML, expected);
```

Коли внутрішній HTML `stringToTest` дорівнює `Gu1n34 P1g5`, значення `regexPattern` дорівнює `G`, і прапорці глобального режиму та нечутливості до регістру увімкнено, `stringToTest.innerHTML` має стати `<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5` при натисканні кнопки `testButton`.

```js
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "G";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">G</span>u1n34 P1<span class="highlight">g</span>5';
assert.equal(stringToTest.innerHTML, expected);
```

Коли ви натискаєте кнопку `testButton`, якщо шаблон регулярного виразу співпадає з тестовим рядком, знайдений текст має бути оточений елементом `span` з класом `highlight`.

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = '<span class="highlight">Sh</span>e <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls <span class="highlight">sea</span><span class="highlight">sh</span>ells by the <span class="highlight">sea</span><span class="highlight">sh</span>ore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = 'She <span class="highlight">se</span>lls seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = '<span class="highlight">Sh</span>e sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

Якщо збігів немає, тестовий рядок не повинен змінюватися.

```js
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = 'She sells seashells by the seashore.';
assert.strictEqual(stringToTest.innerHTML, expected);
```

Коли внутрішній HTML `stringToTest` дорівнює `Gu1n34 P1g5`, значення `regexPattern` дорівнює `\d+`, і жоден прапорець не увімкнено, внутрішній текст `#result` має стати `1` при натисканні кнопки `testButton`.

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = false;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1"
assert.strictEqual(out.innerText, expected)
```

Коли внутрішній HTML `stringToTest` дорівнює `Gu1n34 P1g5`, значення `regexPattern` дорівнює `\d+`, і прапорець глобального режиму увімкнено, внутрішній текст `#result` має стати `1, 34, 1, 5` при натисканні кнопки `testButton`.

```js
const out = document.getElementById("result");
stringToTest.innerText = "Gu1n34 P1g5"
regexPattern.value = "\\d+";
caseInsensitiveFlag.checked = false;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "1, 34, 1, 5"
assert.strictEqual(out.innerText, expected)
```

Якщо є збіги, знайдений текст має відображатися всередині `#result`.

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "s[he]a?";
caseInsensitiveFlag.checked = true;
globalFlag.checked = true;
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "Sh, se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se, sea, sh, sea, sh"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
globalFlag.checked = false;
testBtn.dispatchEvent(new Event("click"));
expected = "se"
assert.strictEqual(out.innerText, expected)

stringToTest.innerText = "She sells seashells by the seashore."
caseInsensitiveFlag.checked = true;
testBtn.dispatchEvent(new Event("click"));
expected = "Sh"
assert.strictEqual(out.innerText, expected)
```

Якщо збігів немає, у `no match` має відображатися текст `#result`.

```js
const out = document.getElementById("result");
stringToTest.innerText = "She sells seashells by the seashore."
regexPattern.value = "sea\\s";
const testBtn = document.getElementById("test-btn");
testBtn.dispatchEvent(new Event("click"));
let expected = "no match";
assert.strictEqual(out.innerText, expected);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
    <title>Regex Sandbox</title>

</head>

<body>
    <h1>Regex Sandbox</h1>
    <main>
        <div id="regex-container">
            <label for="pattern">Regex Pattern:
                <div id="pattern-container">/<input type="text" id="pattern" name="pattern"
                        placeholder="Enter your regex pattern">/</div>
            </label>
            <div id="flags-container">
                <p>Flags: </p>
                <label for="i">
                    <input type="checkbox" name="flags" id="i"> i
                </label>
                <label for="g">
                    <input type="checkbox" name="flags" id="g"> g
                </label>
            </div>
        </div>
        <div id="test-container">
            <p>Test String:</p>
            <div id="test-string" placeholder="Enter your test string" contenteditable="true"></div>
        </div>
        <button class="btn" id="test-btn" type="button">Test Regex</button>
        <div id="result-container">
            <h2>Result:</h2>
            <p id="result">
            </p>
        </div>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

:root {
    --dark-grey: #1b1b32;
    --light-grey: #f5f6f7;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
    --border: 0.2rem solid darkgrey;
    --padding: 0.3rem;
}

body {
    background-color: var(--dark-grey);
    color: var(--light-grey);
    font-size: 20px;
    font-family: "Lato", Helvetica, Arial, sans-serif;
    padding: 5px;
}

h1 {
    margin: 5rem auto 2rem;
    text-align: center;
}

p {
    padding: var(--padding);
}

#regex-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
    border: var(--border);
}

#regex-container>label {
    padding: var(--padding);
    flex: 1 1 auto;
}

#pattern-container {
    display: inline-block;
    color: var(--dark-grey);
    background-color: var(--light-grey);
    margin: 5px;
    border: var(--border);
}

#pattern {
    margin: 0.2rem;
    border: 0;
    font-size: 1rem;
    width: calc(100% - 1.2rem);
}

#pattern:focus {
    outline: none;
}

#flags-container {
    display: flex;
    align-items: center;
    flex: 1 1 auto;
}

#flags-container>label {
    padding: var(--padding);
    margin-right: 0.3rem;
}

#test-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    flex-direction: column;
    flex: 0 0 auto;
    border: var(--border);
}

#test-string {
    background-color: var(--light-grey);
    min-height: 5rem;
    color: var(--dark-grey);
    border-top: var(--border);
    font-size: 1.2rem;
}

[contenteditable=true]:empty:before {
    content: attr(placeholder);
    pointer-events: none;
    color: var(--dark-grey);
}

::placeholder {
    color: var(--dark-grey);
}


button {
    display: block;
    cursor: pointer;
    width: 8rem;
    margin: 0.2rem auto;
    color: var(--dark-grey);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 0.2rem;
    font-size: 1.1rem;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
}

h2 {
    align-self: flex-start;
    margin: 0.4rem 0.2rem 0.2rem;
    flex: 0 1 auto;
}

#result {
    color: var(--dark-grey);
    background-color: var(--light-grey);
    font-size: 1.5rem;
    flex: 1 1 auto;
    margin: 0.2rem;
    border: var(--border);
    min-height: 3rem;
}

.highlight {
    background-color: lightgreen;
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <link rel="stylesheet" href="styles.css">
    <title>Regex Sandbox</title>

</head>

<body>
    <h1>Regex Sandbox</h1>
    <main>
        <div id="regex-container">
            <label for="pattern">Regex Pattern:
                <div id="pattern-container">/<input type="text" id="pattern" name="pattern"
                        placeholder="Enter your regex pattern">/</div>
            </label>
            <div id="flags-container">
                <p>Flags: </p>
                <label for="i">
                    <input type="checkbox" name="flags" id="i"> i
                </label>
                <label for="g">
                    <input type="checkbox" name="flags" id="g"> g
                </label>
            </div>
        </div>
        <div id="test-container">
            <p>Test String:</p>
            <div id="test-string" placeholder="Enter your test string" contenteditable="true"></div>
        </div>
        <button class="btn" id="test-btn" type="button">Test Regex</button>
        <div id="result-container">
            <h2>Result:</h2>
            <p id="result">
            </p>
        </div>

    </main>
    <script src="script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    box-sizing: border-box;
    margin: 0;
    padding: 0;
}

:root {
    --dark-grey: #1b1b32;
    --light-grey: #f5f6f7;
    --golden-yellow: #fecc4c;
    --yellow: #ffcc4c;
    --gold: #feac32;
    --orange: #ffac33;
    --dark-orange: #f89808;
    --border: 0.2rem solid darkgrey;
    --padding: 0.3rem;
}

body {
    background-color: var(--dark-grey);
    color: var(--light-grey);
    font-size: 20px;
    font-family: "Lato", Helvetica, Arial, sans-serif;
}

h1 {
    margin: 5rem auto 2rem;
    text-align: center;
}

p {
    padding: var(--padding);
}

#regex-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
    border: var(--border);
}

#regex-container>label {
    padding: var(--padding);
    flex: 1 1 auto;
}

#pattern-container {
    display: inline-block;
    color: var(--dark-grey);
    background-color: var(--light-grey);
    margin: 5px;
    border: var(--border);
}

#pattern {
    margin: 0.2rem;
    border: 0;
    font-size: 1rem;
    width: calc(100% - 1.2rem);
}

#pattern:focus {
    outline: none;
}

#flags-container {
    display: flex;
    align-items: center;
    flex: 1 1 auto;
}

#flags-container>label {
    padding: var(--padding);
    margin-right: 0.3rem;
}

#test-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    flex-direction: column;
    flex: 0 0 auto;
    border: var(--border);
}

#test-string {
    background-color: var(--light-grey);
    min-height: 5rem;
    color: var(--dark-grey);
    border-top: var(--border);
    font-size: 1.2rem;
}

[contenteditable=true]:empty:before {
    content: attr(placeholder);
    pointer-events: none;
    color: var(--dark-grey);
}

::placeholder {
    color: var(--dark-grey);
}


button {
    display: block;
    cursor: pointer;
    width: 8rem;
    margin: 0.2rem auto;
    color: var(--dark-grey);
    background-color: var(--gold);
    background-image: linear-gradient(var(--golden-yellow), var(--orange));
    border-color: var(--gold);
    border-width: 0.2rem;
    font-size: 1.1rem;
}

.btn:hover {
    background-image: linear-gradient(var(--yellow), var(--dark-orange));
}

#result-container {
    max-width: 680px;
    margin: 20px auto;
    display: flex;
    justify-content: center;
    align-items: center;
}

h2 {
    align-self: flex-start;
    margin: 0.4rem 0.2rem 0.2rem;
    flex: 0 1 auto;
}

#result {
    color: var(--dark-grey);
    background-color: var(--light-grey);
    font-size: 1.5rem;
    flex: 1 1 auto;
    margin: 0.2rem;
    border: var(--border);
    min-height: 3rem;
}

.highlight {
    background-color: lightgreen;
}
```

```js
const regexPattern = document.getElementById("pattern");
const caseInsensitiveFlag = document.getElementById("i");
const globalFlag = document.getElementById("g");
const testResult = document.getElementById("result");
const getFlags = () => {
    let flags = '';
    const flagArray = [caseInsensitiveFlag, globalFlag]
    flagArray.forEach((flag) => {
        if (flag.checked) flags += flag.id;
    })
    return flags
}
const stringToTest = document.getElementById("test-string");
const testButton = document.getElementById("test-btn")
testButton.addEventListener(
    "click",
    () => {
        let matched;
        let flags = getFlags();
        const regex = new RegExp(regexPattern.value, flags)
        if (flags.includes(globalFlag)) {
            stringToTest.innerHTML = stringToTest.innerText.replaceAll(regex, '<span class="highlight">$&</span>')
            matched = stringToTest.innerText.matchAll(regex)
        } else {
            stringToTest.innerHTML = stringToTest.innerText.replace(regex, '<span class="highlight">$&</span>')
            matched = stringToTest.innerText.match(regex);             
        }
        testResult.innerText = matched ? matched.join(", ") : 'no match';
    }
)
```
