---
id: bad87fee1348bd9aec908853
title: id-Attribute zu Bootstrap-Elementen hinzufügen
challengeType: 0
forumTopicId: 16639
dashedName: add-id-attributes-to-bootstrap-elements
---

# --description--

Erinnern Sie sich, dass Sie neben Klassenattributen jedem Ihrer Elemente ein `id`-Attribut geben können.

Jede id muss eindeutig für ein bestimmtes Element sein und darf pro Seite nur einmal verwendet werden.

Geben wir jedem unserer `div`-Elemente der Klasse `well` eine eindeutige id.

Denken Sie daran, dass Sie einem Element eine id so geben können:

```html
<div class="well" id="center-well">
```

Geben Sie dem Well auf der linken Seite die id `left-well`. Geben Sie dem Well auf der rechten Seite die id `right-well`.

# --hints--

Ihr linkes `well` sollte die id `left-well` haben.

```js
const column = document.querySelectorAll('.col-xs-6')[0];
const leftWall = column?.querySelectorAll(`:scope ${'#left-well'}`);
assert.lengthOf(leftWall,1)
```

Ihr rechtes `well` sollte die id `right-well` haben.

```js
const column = document.querySelectorAll('.col-xs-6')[1];
const rightWall = column?.querySelectorAll(`:scope ${'#right-well'}`);
assert.lengthOf(rightWall,1)
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well" id="left-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well" id="right-well">
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
        <button class="btn btn-default target"></button>
      </div>
    </div>
  </div>
</div>
```
