---
id: bad87fee1348bd9aec908846
title: Erstellen Sie eine Bootstrap-Überschrift
challengeType: 0
forumTopicId: 16812
dashedName: create-a-bootstrap-headline
---

# --description--

Jetzt bauen wir etwas von Grund auf, um unsere HTML-, CSS- und Bootstrap-Fähigkeiten zu üben.

Wir erstellen einen jQuery-Spielplatz, den wir bald in unseren jQuery-Aufgaben verwenden werden.

Beginnen Sie damit, ein `h3`-Element mit dem Text `jQuery Playground` zu erstellen.

Färben Sie Ihr `h3`-Element mit der Bootstrap-Klasse `text-primary` und zentrieren Sie es mit der Bootstrap-Klasse `text-center`.

# --hints--

Sie sollten Ihrer Seite ein `h3`-Element hinzufügen.

```js
assert.lengthOf(document.querySelectorAll('h3'),1);
```

Ihr `h3`-Element sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/h3>/g);
assert.match(code,/<h3/g);
assert.equal( code.match(/<\/h3>/g).length , code.match(/<h3/g).length);
```

Ihr `h3`-Element sollte durch Anwenden der Klasse `text-primary` eingefärbt werden.

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-primary'));
```

Ihr `h3`-Element sollte durch Anwenden der Klasse `text-center` zentriert werden.

```js
assert.isTrue(document.querySelector('h3')?.classList?.contains('text-center'));
```

Ihr `h3`-Element sollte den Text `jQuery Playground` enthalten.

```js
assert.match(document.querySelector('h3')?.textContent, /jquery(\s)+playground/gi);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```
