---
id: bad87fee1348bd9aec908849
title: Elemente innerhalb Ihrer Bootstrap-Wells hinzufügen
challengeType: 0
forumTopicId: 16636
dashedName: add-elements-within-your-bootstrap-wells
---

# --description--

Jetzt sind wir mehrere `div`-Elemente tief in jeder Spalte unserer Zeile. So tief müssen wir gehen. Nun können wir unsere `button`-Elemente hinzufügen.

Verschachteln Sie drei `button`-Elemente innerhalb jedes Ihrer `div`-Elemente mit dem Klassennamen `well`.

# --hints--

Drei `button`-Elemente sollten innerhalb jedes Ihrer `div`-Elemente mit der Klasse `well` verschachtelt sein.

```js
const buttonOne = document.querySelectorAll('div.well')?.[0];
const buttonOneChildren = buttonOne?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonOneChildren,3);

const buttonTwo = document.querySelectorAll('div.well')?.[1]
const buttonTwoChildren = buttonTwo?.querySelectorAll(`:scope ${'button'}`);
assert.lengthOf(buttonTwoChildren,3);
```

Sie sollten insgesamt 6 `button`-Elemente haben.

```js
assert.lengthOf(document.querySelectorAll('button'), 6);
```

Alle Ihre `button`-Elemente sollten schließende Tags haben.

```js
assert.match(code,/<\/button>/g);
assert.match(code,/<button/g);
assert.equal(code.match(/<\/button>/g).length,code.match(/<button/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">



      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
    <div class="col-xs-6">
      <div class="well">
        <button></button>
        <button></button>
        <button></button>
      </div>
    </div>
  </div>
</div>
```
