---
id: bad87fee1348bd9aec908847
title: Teilen Sie Ihre Bootstrap-Zeile
challengeType: 0
forumTopicId: 18306
dashedName: split-your-bootstrap-row
---

# --description--

Da wir jetzt eine Bootstrap-Zeile haben, teilen wir sie in zwei Spalten auf, um unsere Elemente unterzubringen.

Erstellen Sie zwei `div`-Elemente innerhalb Ihrer Zeile, beide mit der Klasse `col-xs-6`.

# --hints--

Zwei `div class="col-xs-6"`-Elemente sollten innerhalb Ihres `div class="row"`-Elements verschachtelt sein.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Alle Ihre `div`-Elemente sollten schließende Tags haben.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g).length , code.match(/<div/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">


  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6"></div>
    <div class="col-xs-6"></div>
  </div>
</div>
```
