---
id: bad87fee1348bd9aec908857
title: "Kommentare verwenden, um Code zu erläutern"
challengeType: 0
forumTopicId: 18347
dashedName: use-comments-to-clarify-code
---

# --description--

Wenn wir anfangen, jQuery zu verwenden, werden wir HTML-Elemente ändern, ohne sie tatsächlich im HTML anzupassen.

Stellen Sie sicher, dass alle wissen, dass sie diesen Code nicht direkt verändern sollen.

Denken Sie daran, dass Sie einen Kommentar mit `<!--` beginnen und mit `-->` beenden können.

Fügen Sie oben in Ihrem HTML einen Kommentar hinzu, der `Code below this line should not be changed` sagt.

# --hints--

Sie sollten einen Kommentar mit `<!--` oben in Ihrem HTML beginnen.

```js
assert.match(code,(/^\s*<!--/));
```

Ihr Kommentar sollte den Text `Code below this line should not be changed` enthalten.

```js
assert.match(code,(/<!--(?!(>|->|.*-->.*this line))\s*.*this line.*\s*-->/gi));
```

Sie sollten Ihren Kommentar mit `-->` schließen.

```js
assert.match(code,(/-->.*\n+.+/g));
```

Sie sollten die gleiche Anzahl von Kommentar-Anfängern und -Endern haben.

```js
assert.equal(code.match(/<!--/g).length,code.match(/-->/g).length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```

# --solutions--

```html
<!-- Code below this line should not be changed -->
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <h4>#left-well</h4>
      <div class="well" id="left-well">
        <button class="btn btn-default target" id="target1">#target1</button>
        <button class="btn btn-default target" id="target2">#target2</button>
        <button class="btn btn-default target" id="target3">#target3</button>
      </div>
    </div>
    <div class="col-xs-6">
      <h4>#right-well</h4>
      <div class="well" id="right-well">
        <button class="btn btn-default target" id="target4">#target4</button>
        <button class="btn btn-default target" id="target5">#target5</button>
        <button class="btn btn-default target" id="target6">#target6</button>
      </div>
    </div>
  </div>
</div>
```
