---
id: bad87fee1348bd9aec908746
title: Unsere Seite in einem Bootstrap container-fluid div unterbringen
challengeType: 0
forumTopicId: 18198
dashedName: house-our-page-within-a-bootstrap-container-fluid-div
---

# --description--

Stellen Sie nun sicher, dass alle Inhalte auf Ihrer Seite mobil responsiv sind.

Verschachteln Sie Ihr `h3`-Element innerhalb eines `div`-Elements mit der Klasse `container-fluid`.

# --hints--

Ihr `div`-Element sollte die Klasse `container-fluid` haben.

```js
assert.isTrue(document.querySelector('div')?.classList?.contains('container-fluid'));
```

Jedes Ihrer `div`-Elemente sollte ein schließendes Tag haben.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);

assert.equal(code.match(/<\/div>/g).length ,code.match(/<div/g).length);
```

Ihr `h3`-Element sollte innerhalb eines `div`-Elements verschachtelt sein.

```js
const divElement = document.querySelector('div');
const divChildren = divElement?.querySelectorAll(`:scope ${'h3'}`)
assert.lengthOf(divChildren,1);
```

# --seed--

## --seed-contents--

```html
<h3 class="text-primary text-center">jQuery Playground</h3>
```

# --solutions--

```html
<div class="container-fluid">
    <h3 class="text-primary text-center">jQuery Playground</h3>
</div>
```
