---
id: bad87fee1348bd9aec908848
title: Bootstrap-Wells erstellen
challengeType: 0
forumTopicId: 16825
dashedName: create-bootstrap-wells
---

# --description--

Bootstrap hat eine Klasse namens `well`, die für Ihre Spalten einen visuellen Tiefeneffekt erzeugen kann.

Verschachteln Sie jeweils ein `div`-Element mit der Klasse `well` innerhalb jedes Ihrer `col-xs-6` `div`-Elemente.

# --hints--

Sie sollten ein `div`-Element mit der Klasse `well` in jedes Ihrer `div`-Elemente mit der Klasse `col-xs-6` einfügen.

```js
const wells = document.querySelectorAll('div.col-xs-6 > div.well');
assert.lengthOf( wells,2 ); 
```

Beide Ihrer `div`-Elemente mit der Klasse `col-xs-6` sollten innerhalb Ihres `div`-Elements mit der Klasse `row` verschachtelt sein.

```js
assert.lengthOf(document.querySelectorAll('div.row > div.col-xs-6'),2);
```

Alle Ihre `div`-Elemente sollten schließende Tags haben.

```js
assert.match(code,/<\/div>/g);
assert.match(code,/<div/g);
assert.equal(code.match(/<\/div>/g)?.length , code.match(/<div/g)?.length);
```

# --seed--

## --seed-contents--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">

    </div>
    <div class="col-xs-6">

    </div>
  </div>
</div>
```

# --solutions--

```html
<div class="container-fluid">
  <h3 class="text-primary text-center">jQuery Playground</h3>
  <div class="row">
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
    <div class="col-xs-6">
      <div class="well"></div>
    </div>
  </div>
</div>
```
