---
id: 69a56b5069ca99f7317e6e19
title: Jenga Kagua Manifesti ya Mizigo
challengeType: 26
dashedName: lab-cargo-manifest-validator
---

# --description--

Katika maabara hii, utatumia JavaScript kawaisha na kukagua manifesti za mizigo. Manifesti ya mizigo ni hati inayoorodhesha bidhaa zinazobebwa (kwa mfano, kwa meli au treni) na ina maelezo kuhusu bidhaa hizo.

Kila manifesti ya mizigo itawakilishwa kama kitu chenye sifa zifuatazo:

- `containerId`: nambari kamili chanya inayotambulisha kontena la mizigo linalohusiana.
- `destination`: mfuatano wa herufi usio tupu (baada ya kuondoa nafasi zisizohitajika) unaoonyesha mahali pa mwisho pa mizigo.
- `weight`: nambari chanya inayowakilisha uzito wa mizigo.
- `unit`: mfuatano wa herufi unaoelezea kitengo cha sifa ya `weight` ya mizigo (iwe ni `"kg"` kwa kilo au `"lb"` kwa pauni).
- `hazmat`: thamani ya kweli au siyo kweli inayoonyesha kama kuna hitaji la kushughulikia vifaa hatari.

Mfano wa kitu cha manifesti ya mizigo:

```js
{
  containerId: 1,
  destination: "Monterey, California, USA",
  weight: 831,
  unit: "lb",
  hazmat: false
}
```

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kutekeleza kitendakazi kinachoitwa `normalizeUnits` chenye kigezo `manifest`.
    - Kitendakazi hakipaswi kubadilisha kontena la asili la manifest na lazima kila mara kirudishe kontena jipya ambapo `weight` imekuwa kawaishwa kwa kilo na `unit` imewekwa kuwa `"kg"`.
    - Ikiwa `weight` ya kontena la manifest imeelezwa kwa pauni (`unit: "lb"`), kitendakazi kinapaswa kubadilisha `weight` kuwa kilo kwa kutumia mabadiliko ya takriban `1 lb = 0.45 kg`, na kusahihisha `unit` ipasavyo.
    - Ikiwa `weight` tayari imeelezwa kwa kilo (`unit: "kg"`), `weight` na `unit` zinapaswa kubaki kama zilivyo.
2. Unapaswa kutekeleza kitendakazi kinachoitwa `validateManifest` chenye kigezo `manifest`.
    - Kitendakazi hakipaswi kubadilisha kontena la asili la manifest na lazima kila mara kirudishe kontena jipya.
    - Ikiwa manifest ya ingizo ni halali (hakuna sifa zilizokosekana au batili), kitendakazi kinapaswa kurudisha kontena tupu.
    - Ikiwa manifest ya ingizo si halali, kitendakazi kinapaswa kurudisha kontena chenye maingizo kwa kila sifa iliyokosekana au batili. Sifa zilizokosekana zinapaswa kuwa na thamani `"Missing"` na sifa batili ziwe na thamani `"Invalid"`.

    Mfano wa thamani ya return ambapo kitu cha ingizo hakina sifa ya `destination` na kina sifa batili ya `weight`:

    ```js
    {
      destination: "Missing",
      weight: "Invalid"
    }
    ```

3. Unapaswa kutekeleza kitendakazi kinachoitwa `processManifest` chenye kigezo `manifest`. Kitendakazi hiki kinapaswa kuandika kwenye kumbukumbu:
    - Ikiwa kitu cha manifest ni halali, andika `Validation success: ${containerId}` kisha uzingatie `weight` ya manifest kwa kilo, kwa muundo wa `Total weight: ${weight} kg`. Tumia `normalizeUnits()` kwa ajili ya uongofu huu.
    - Ikiwa kitu cha manifest si halali, andika `Validation error: ${containerId}` kisha andika kitu kilichorejeshwa na wito wa `validateManifest()` ukiwa na kitu cha manifest.

   **Kumbuka:** Kila mojawapo ya hali hizi mbili inapaswa kuwa na wito mbili za `console.log()`.

**Kumbuka:** Usitangaze `normalizeUnits`, `validateManifest`, au `processManifest` kwa kutumia `const`, kwa sababu majaribio yanahitaji kuyarekebisha.

# --before-each--

```js
const _validManifests = [
  {
    containerId: 4,
    destination: "Monterey, California, USA",
    weight: 831,
    unit: "lb",
    hazmat: false
  },
  {
    containerId: 5,
    destination: "Montreal, Quebec, Canada",
    weight: 151,
    unit: "kg",
    hazmat: false
  }
];

const _invalidManifests = [
  {
    containerId: -6,
    destination: 123,
    weight: 0,
    unit: "pounds",
    hazmat: true
  },
  {
    containerId: 0,
    weight: -21,
    unit: "KG"
  },
  {
    containerId: "eight",
    destination: "Guadalajara, Jalisco, Mexico",
    weight: 9001,
    hazmat: "no"
  },
  {
    destination: " ",
    weight: NaN,
    unit: "kg",
    hazmat: 1
  },
  {
    containerId: 10,
    destination: " ",
  },
];

const _testObj = {
  containerId: 1.51,
  destination: "Lavender Town",
  weight: NaN,
};

console.log = () => {};
```

# --hints--

Hupaswi kutangaza `normalizeUnits`, `validateManifest`, au `processManifest` kwa kutumia `const`, kwa sababu majaribio yanahitaji kuyarekebisha.

```js
assert.doesNotThrow(() => {
  normalizeUnits = normalizeUnits;
  validateManifest = validateManifest;
  processManifest = processManifest;
});
```

Unapaswa kuwa na kitendakazi kinachoitwa `normalizeUnits` chenye kigezo `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.normalizeUnits);
const params = explorer.allFunctions.normalizeUnits.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Kitendakazi chako cha `normalizeUnits` kinapaswa kurudisha kontena jipya `{ containerId: 68, destination: "Salinas", weight: 45.45, unit: "kg", hazmat: true }` kinapoitwa na `{ containerId: 68, destination: "Salinas", weight: 101, unit: "lb", hazmat: true }`, bila kubadilisha kontena la asili.

```js
const testObj = {
  containerId: 68,
  destination: "Salinas",
  weight: 101,
  unit: "lb",
  hazmat: true
};

const expected = {
  containerId: 68,
  destination: "Salinas",
  weight: 45.45,
  unit: "kg",
  hazmat: true
};

const copy = normalizeUnits(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "normalizeUnits() did not return the expected normalized object.");
assert.notStrictEqual(copy, testObj, "normalizeUnits() should return a new object.");
assert.strictEqual(testObj.containerId, 68, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Salinas", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 101, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "lb", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, true, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 5, "Original object should not be mutated.");
```

Kitendakazi chako cha `normalizeUnits` kinapaswa kurudisha nakala ya kontena la manifest la ingizo na `weight` yake kawaishwa kwa kilo na `unit` yake kuwekwa kuwa `"kg"`. Tumia mabadiliko ya takriban `1 lb = 0.45 kg` kwa mabadiliko ya `weight`.

```js
_validManifests.forEach((obj, i) => {
  const result = normalizeUnits(obj);
  assert.isObject(result, "normalizeUnits() should return an object.");
  assert.notStrictEqual(result, obj, "normalizeUnits() should return a new object, not mutate the original.");
  const expectedWeight = obj.unit === "lb" ? obj.weight * 0.45 : obj.weight;
  const expected = {
    ...obj,
    weight: expectedWeight,
    unit: "kg"
  };

  assert.deepEqual(
    result,
    expected,
    `normalizeUnits() did not return the expected normalized object for manifest at index ${i}.`
  );
});
```

Kitendakazi chako cha `normalizeUnits` kinapaswa kurudisha nakala mpya ya kontena la manifest la ingizo bila kubadilisha kontena la asili.

```js
const original = {
  containerId: 151,
  destination: "Saffron City",
  weight: 151,
  unit: "lb",
  hazmat: false
};

const expected = {
  containerId: 151,
  destination: "Saffron City",
  weight: 67.95,
  unit: "kg",
  hazmat: false
};

const copy = normalizeUnits(original);
assert.isObject(copy);
assert.deepEqual(copy, expected);
assert.notStrictEqual(copy, original, "normalizeUnits() should return a new object.");
assert.strictEqual(original.weight, 151, "Original object should not be mutated.");
assert.strictEqual(original.unit, "lb", "Original object should not be mutated.");
assert.equal(Object.keys(original).length, 5, "Original object should not be mutated.");
```

Unapaswa kuwa na kitendakazi kinachoitwa `validateManifest` chenye kigezo `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.validateManifest);
const params = explorer.allFunctions.validateManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya, tupu `{}` kinapoitwa na `{ containerId: 1, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`.

```js
const testObj = {
  containerId: 1,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for a valid manifest.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 1, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, "Santa Cruz", "Original object should not be mutated.");
assert.strictEqual(testObj.weight, 304, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, false, "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena tupu `{}` ikiwa kontena la manifest la ingizo ni halali.

```js
const solution = {};
for (const obj of _validManifests) {
  assert.deepEqual(
    validateManifest(obj),
    solution,
    "validateManifest() should return an empty object for valid manifests."
  );
}
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Missing", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{}`, bila kubadilisha kontena la asili.

```js
const testObj = {};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.equal(Object.keys(testObj).length, 0, "validateManifest() should return a new object, not mutate the original.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Invalid" }` kinapoitwa na `{ containerId: null, destination: "Santa Cruz", weight: 304, unit: "kg", hazmat: false }`, bila kubadilisha kontena la asili.

```js
const testObj = {
  containerId: null,
  destination: "Santa Cruz",
  weight: 304,
  unit: "kg",
  hazmat: false
};

const expected = {
  containerId: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for null containerId.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Invalid", destination: "Invalid", weight: "Invalid", unit: "Invalid", hazmat: "Invalid" }` kinapoitwa na `{ containerId: 0, destination: 405, weight: -84, unit: "pounds", hazmat: "no" }`, bila kubadilisha kontena la asili.

```js
const testObj = {
  containerId: 0,
  destination: 405,
  weight: -84,
  unit: "pounds",
  hazmat: "no"
};

const expected = {
  containerId: "Invalid",
  destination: "Invalid",
  weight: "Invalid",
  unit: "Invalid",
  hazmat: "Invalid"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 0, "Original object should not be mutated.");
assert.strictEqual(testObj.destination, 405, "Original object should not be mutated.");
assert.strictEqual(testObj.weight, -84, "Original object should not be mutated.");
assert.strictEqual(testObj.unit, "pounds", "Original object should not be mutated.");
assert.strictEqual(testObj.hazmat, "no", "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{ containerId: -2 }`, bila kubadilisha kontena la asili.

```js
const testObj = {
  containerId: -2,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, -2, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Invalid", destination: "Missing", weight: "Missing", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{ containerId: 3.50 }`, bila kubadilisha kontena la asili. Unaweza kutumia `Number.isInteger()` kuthibitisha thamani za nambari kamili.

```js
const testObj = {
  containerId: 3.50,
};

const expected = {
  containerId: "Invalid",
  destination: "Missing",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.containerId, 3.50, "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Missing", destination: "Invalid", weight: "Missing", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{ destination: "  " }`, bila kubadilisha kontena la asili. Unaweza kutumia `.trim()` kuondoa nafasi zisizohitajika kutoka kwenye mfuatano wa herufi.

```js
const testObj = {
  destination: "  ",
};

const expected = {
  containerId: "Missing",
  destination: "Invalid",
  weight: "Missing",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.strictEqual(testObj.destination, "  ", "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena jipya `{ containerId: "Missing", destination: "Missing", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{ weight: NaN }`, bila kubadilisha kontena la asili. Unaweza kutumia `Number.isNaN()` kuthibitisha thamani za `NaN`.

```js
const testObj = {
  weight: NaN,
};

const expected = {
  containerId: "Missing",
  destination: "Missing",
  weight: "Invalid",
  unit: "Missing",
  hazmat: "Missing"
};

const copy = validateManifest(testObj);
assert.isObject(copy);
assert.deepEqual(copy, expected, "validateManifest() did not return the expected object for invalid and missing properties.");
assert.notStrictEqual(copy, testObj, "validateManifest() should return a new object, not mutate the original.");
assert.ok(Number.isNaN(testObj.weight), "Original object should not be mutated.");
assert.equal(Object.keys(testObj).length, 1, "Original object should not be mutated.");
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kontena kinachoelezea sifa zilizokosekana na/au batili ikiwa kontena la manifest la ingizo si halali.

```js
const solutions = [
  { containerId: 'Invalid', destination: 'Invalid', weight: 'Invalid', unit: 'Invalid' },
  { containerId: 'Invalid', destination: 'Missing', weight: 'Invalid', unit: 'Invalid', hazmat: 'Missing' },
  { containerId: 'Invalid', unit: 'Missing', hazmat: 'Invalid' },
  { containerId: 'Missing', destination: 'Invalid', weight: 'Invalid', hazmat: 'Invalid' },
  { destination: 'Invalid', weight: 'Missing', unit: 'Missing', hazmat: 'Missing' }
];

let i = 0;
for (const obj of _invalidManifests) {
  const result = validateManifest(obj);
  assert.deepEqual(
    result,
    solutions[i],
    "validateManifest() should return an object with errors for invalid manifests."
  );
  i += 1;
}
```

Kitendakazi chako cha `validateManifest` kinapaswa kurudisha kitu kipya bila kubadilisha asili.

```js
const invalidOriginal = {
  containerId: -151,
  destination: "",
  weight: -151,
  unit: "kilograms",
  hazmat: "idk"
};

const validOriginal = {
  containerId: 151,
  destination: "Pallet Town",
  weight: 151,
  unit: "kg",
  hazmat: false
};

const res0 = validateManifest(invalidOriginal);
assert.isObject(res0);
assert.notStrictEqual(res0, invalidOriginal, "validateManifest() should return a new object.");
assert.strictEqual(invalidOriginal.containerId, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.destination, "", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.weight, -151, "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.unit, "kilograms", "Original object should not be mutated.");
assert.strictEqual(invalidOriginal.hazmat, "idk", "Original object should not be mutated.");
assert.equal(Object.keys(invalidOriginal).length, 5, "Original object should not be mutated.");

const res1 = validateManifest(validOriginal);
assert.isObject(res1);
assert.notStrictEqual(res1, validOriginal, "validateManifest() should return a new object.");
assert.strictEqual(validOriginal.containerId, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.destination, "Pallet Town", "Original object should not be mutated.");
assert.strictEqual(validOriginal.weight, 151, "Original object should not be mutated.");
assert.strictEqual(validOriginal.unit, "kg", "Original object should not be mutated.");
assert.strictEqual(validOriginal.hazmat, false, "Original object should not be mutated.");
assert.equal(Object.keys(validOriginal).length, 5, "Original object should not be mutated.");
```

Unapaswa kuwa na kitendakazi kinachoitwa `processManifest` chenye kigezo `manifest`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.processManifest);
const params = explorer.allFunctions.processManifest.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), 'manifest');
```

Kitendakazi chako cha `processManifest` kinapaswa kwanza kuandika `Validation success: 55` kisha kuandika `Total weight: 180 kg` kinapoitwa na `{ containerId: 55, destination: "Carmel", weight: 400, unit: "lb", hazmat: false }`.

```js
const testObj = {
  containerId: 55,
  destination: "Carmel",
  weight: 400,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);

  const expectedCalls = [
    ["Validation success: 55"],
    ["Total weight: 180 kg"]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation success message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the normalized weight in kg.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kwanza kuandika ujumbe wa mafanikio, `Validation success: ${containerId}`, ikiwa kontena la manifest la ingizo ni halali.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kawaisha kontena la manifest la ingizo lililo halali kuwa kilo kwa kutumia `normalizeUnits()` kisha kuandika `Total weight: ${weight} kg`.

```js
const testObj = {
  containerId: 97,
  destination: "Cerulean City",
  weight: 200,
  unit: "lb",
  hazmat: false
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation success: 97"],
    ["Total weight: 90 kg"],
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kuchapisha ujumbe wa mafanikio ukiwa na `containerId` ya kitu, kisha kichapishe `weight` ya kitu hicho kwa kilogramu, ikiwa ingizo la manifest ni sahihi. Unapaswa kutumia `normalizeUnits()` kwa ajili ya uongofu huu, na hii inapaswa kujumuisha miito miwili ya `console.log()`.

```js
const originalNormalize = normalizeUnits;
let called = false;
try {
  normalizeUnits = function (...args) {
    called = true;
    return originalNormalize(...args);
  };

  processManifest(_validManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call normalizeUnits() when the manifest is valid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  normalizeUnits = originalNormalize;
}

const spy = __helpers.spyOn(console, "log");

try {
  for (const obj of _validManifests) {
    processManifest(obj);
  }
  const expectedCalls = [
    ["Validation success: 4"],
    ["Total weight: 373.95 kg"],
    ["Validation success: 5"],
    ["Total weight: 151 kg"],
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kwanza kuandika `Validation error: -88` kisha kuandika kontena `{ containerId: "Invalid", weight: "Invalid", unit: "Missing", hazmat: "Missing" }` kinapoitwa na `{ containerId: -88, destination: "Soledad", weight: NaN }`.

```js
const testObj = {
  containerId: -88,
  destination: "Soledad",
  weight: NaN
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: -88"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kwanza kuandika `Validation error: undefined` kisha kuandika kontena `{ containerId: "Missing", weight: "Missing", unit: "Missing" }` kinapoitwa na `{ destination: "Watsonville", hazmat: true }`.

```js
const testObj = {
  destination: "Watsonville",
  hazmat: true
};

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(testObj);
  const expectedCalls = [
    ["Validation error: undefined"],
    [{
        containerId: "Missing",
        weight: "Missing",
        unit: "Missing"
    }]
  ];

  assert.deepEqual(spy.calls[0], expectedCalls[0], "First log should be the validation error message.");
  assert.deepEqual(spy.calls[1], expectedCalls[1], "Second log should be the validation result object.");

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kwanza kuandika ujumbe wa kosa, `Validation error: ${containerId}`, ikiwa kontena la manifest la ingizo si halali.

```js
const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[0], expectedCalls[0]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa pia kuandika kontena lililorudishwa na wito wa `validateManifest()` kwa kontena la manifest la asili ikiwa kontena la manifest la ingizo si halali. Piga `console.log()` moja kwa moja na kontena lililorudishwa.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls[1], expectedCalls[1]);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

Kitendakazi chako cha `processManifest` kinapaswa kuandika ujumbe wa kosa na `containerId` ya kontena, kisha kuandika kontena lililorudishwa na wito wa `validateManifest()` kwa kontena la ingizo, ikiwa kontena la manifest la ingizo si halali. Hii inapaswa kuhusisha wito mbili za `console.log()`.

```js
const originalValidate = validateManifest;
let called = false;

try {
  validateManifest = function (...args) {
    called = true;
    return originalValidate(...args);
  };

  processManifest(_invalidManifests[0]);

  assert.isTrue(
    called,
    "processManifest() should call validateManifest() when the manifest is invalid."
  );

} catch (err) {
  assert.fail(err);
} finally {
  validateManifest = originalValidate;
}

const spy = __helpers.spyOn(console, "log");

try {
  processManifest(_testObj);
  const expectedCalls = [
    ["Validation error: 1.51"],
    [{
        containerId: "Invalid",
        weight: "Invalid",
        unit: "Missing",
        hazmat: "Missing"
    }]
  ];
  assert.deepEqual(spy.calls, expectedCalls);

} catch (err) {
  assert.fail(err);
} finally {
  spy.restore();
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function normalizeUnits(manifest) {
  const normalized = {
    ...manifest
  }
  if (normalized.unit === "lb") {
    normalized.weight = normalized.weight * 0.45;
    normalized.unit = "kg";
  }
  return normalized;
}

function validateManifest(manifest) {
  const result = {}

  if (manifest.containerId === undefined) {
    result.containerId = "Missing";
  } else if (
    typeof manifest.containerId !== "number" ||
    !Number.isInteger(manifest.containerId) ||
    manifest.containerId <= 0
  ) {
    result.containerId = "Invalid";
  }
  if (manifest.destination === undefined) {
    result.destination = "Missing";
  } else if (typeof manifest.destination !== "string" || manifest.destination.trim() === "") {
    result.destination = "Invalid";
  }
  if (manifest.weight === undefined) {
    result.weight = "Missing";
  } else if (typeof manifest.weight !== "number" || manifest.weight <= 0 || Number.isNaN(manifest.weight)) {
    result.weight = "Invalid";
  }
  if (manifest.unit === undefined) {
    result.unit = "Missing";
  } else if (manifest.unit !== "kg" && manifest.unit !== "lb") {
    result.unit = "Invalid";
  }
  if (manifest.hazmat === undefined) {
    result.hazmat = "Missing";
  } else if (typeof manifest.hazmat !== "boolean") {
    result.hazmat = "Invalid";
  }

  return result;
}

function processManifest(manifest) {
  const validation = validateManifest(manifest);

  if (Object.keys(validation).length > 0) {
    console.log(`Validation error: ${manifest.containerId}`);
    console.log(validation);
  } else {
    const normalized = normalizeUnits(manifest);
    console.log(`Validation success: ${manifest.containerId}`);
    console.log(`Total weight: ${normalized.weight} kg`);
  }
}
```
