---
id: a77dbc43c33f39daa4429b4f
title: Construir uma Função de Verificação Booleana
challengeType: 26
dashedName: build-a-boolean-check-function
---

# --description--

Neste laboratório, você vai construir uma função que verifica se um valor é classificado como um primitivo booleano.

Os booleanos primitivos são `true` e `false`.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve ter uma função chamada `booWho` que recebe um argumento.
1. Se o argumento recebido for um primitivo booleano, a função deve retornar `true`.
1. Se o argumento for qualquer outro valor, a função deve retornar `false`.

# --hints--

Você deve ter uma função `booWho`.

```js
assert.isFunction(booWho);
```

`booWho(true)` deve retornar `true`.

```js
assert.isTrue(booWho(true));
```

`booWho(false)` deve retornar `true`.

```js
assert.isTrue(booWho(false));
```

`booWho([1,2,3])` deve retornar `false`.

```js
assert.isFalse(booWho([1, 2, 3]));
```

`booWho([].slice)` deve retornar `false`.

```js
assert.isFalse(booWho([].slice));
```

`booWho({"a": 1})` deve retornar `false`.

```js
assert.isFalse(booWho({ a: 1 }));
```

`booWho(1)` deve retornar `false`.

```js
assert.isFalse(booWho(1));
```

`booWho(NaN)` deve retornar `false`.

```js
assert.isFalse(booWho(NaN));
```

`booWho("a")` deve retornar `false`.

```js
assert.isFalse(booWho('a'));
```

`booWho("true")` deve retornar `false`.

```js
assert.isFalse(booWho('true'));
```

`booWho("false")` deve retornar `false`.

```js
assert.isFalse(booWho('false'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function booWho(bool) {
  return typeof bool === 'boolean';
}

```
