---
id: 5d822fd413a79914d39e98ce
title: Step 3
challengeType: 0
dashedName: step-3
---

# --description--

Pia ongeza `box-sizing` ya `border-box` kwa kila kitu. Hii itafanya mpaka uliouongeza usiongeze ukubwa wowote kwa vipengele vyako.

# --hints--

Unapaswa kutumia sifa ya `box-sizing`.

```js
assert.exists(new __helpers.CSSHelp(document).isPropertyUsed('box-sizing'));
```

Unapaswa kutumia kichaguzi kilichopo cha `*`.

```js
// Two selectors create two CSSStyleRule objects
assert.lengthOf(new __helpers.CSSHelp(document).getStyleDeclarations('*'), 1);
```

Vipengele vyote vinapaswa kuwa na `box-sizing` ya `border-box`.

```js
const astStyles = new __helpers.CSSHelp(document).getStyle('*');
assert.equal(astStyles.boxSizing, 'border-box');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
--fcc-editable-region--
* {
  border: 1px solid black;
}

--fcc-editable-region--

```

