---
id: 5d822fd413a79914d39e98d0
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Unda kipengele cha `div` ndani ya `body` chenye darasa la vitu la `background-buildings`. Hii itakuwa kontena la kundi la majengo.

# --hints--

Unapaswa kuunda kipengele cha `div`.

```js
assert.exists(document.querySelector('div'));
```

Kipengele chako cha `div` kinapaswa kuwa ndani ya `body`.

```js
assert.equal(document.querySelector('div')?.parentElement?.localName, 'body');
```

Kipengele chako cha `div` kinapaswa kuwa na darasa la vitu la `background-buildings`

```js
assert.isTrue([...document.querySelector('div')?.classList]?.includes('background-buildings'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

```

