---
id: 5d822fd413a79914d39e98d4
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

Toa sehemu za jengo lako `width` na `height` vigezo na thamani hizi: `70%` na `10%` hadi `.bb1a`, `80%` na `10%` hadi `.bb1b`, `90%` na `10%` hadi `.bb1c`, na `100%` na `70%` hadi `.bb1d`. Kumbuka kuwa asilimia hizi ni kulinganisha na mzazi na pia angalia kuwa urefu utajumlisha hadi 100% - kujaza kontena kwa wima.

# --hints--

Unapaswa kutumia kichaguzi cha darasa la vitu kuunda mtindo wa `.bb1a`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a'));
```

Unapaswa kumpa `.bb1a` `width` ya `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.width, '70%');
```

Unapaswa kumpa `.bb1a` `height` ya `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.height, '10%');
```

Unapaswa kutumia kichaguzi cha darasa la vitu kuunda mtindo wa `.bb1b`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1b'));
```

Unapaswa kumpa `.bb1b` `width` ya `80%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.width, '80%');
```

Unapaswa kumpa `.bb1b` `height` ya `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1b')?.height, '10%');
```

Unapaswa kutumia kichaguzi cha darasa la vitu kuunda mtindo wa `.bb1c`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1c'));
```

Unapaswa kumpa `.bb1c` `width` ya `90%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.width, '90%');
```

Unapaswa kumpa `.bb1c` `height` ya `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1c')?.height, '10%');
```

Unapaswa kutumia kichaguzi cha darasa la vitu kuunda mtindo wa `.bb1d`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1d'));
```

Unapaswa kumpa `.bb1d` `width` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.width, '100%');
```

Unapaswa kumpa `.bb1d` `height` ya `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1d')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
}
--fcc-editable-region--

--fcc-editable-region--

```

