---
id: 5d822fd413a79914d39e98d7
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Ili kutumia variable, weka jina la variable ndani ya mabano ya kawaida ukiwa na `var` mbele yake kama hivi: `var(--variable-name)`. Thamani yoyote uliyoiweka kwa variable itatumika kwa sifa yoyote utakayotumia.

Ongeza variable `--building-color1` uliyotengeneza katika hatua iliyopita kama thamani ya sifa ya `background-color` ya darasa la vitu la `.bb1a`.

# --hints--

Sifa ya `background-color` ya kipengele cha `.bb1a` inapaswa kuwekwa.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1a')?.backgroundColor)
```

Unapaswa kutumia `var(--building-color1)` kuweka sifa ya `background-color` ya kipengele cha `.bb1a`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1a')?.getPropVal('background-color', true), 'var(--building-color1)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />   
  </head>

  <body>
    <div class="background-buildings">
      <div class="bb1">
        <div class="bb1a"></div>
        <div class="bb1b"></div>
        <div class="bb1c"></div>
        <div class="bb1d"></div>
      </div>
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}

.bb1 {
  width: 10%;
  height: 70%;
  display: flex;
  flex-direction: column;
  align-items: center;
  --building-color1: #999;
}

.bb1a {
  width: 70%;
  height: 10%;
--fcc-editable-region--
  
--fcc-editable-region--
}

.bb1b {
  width: 80%;
  height: 10%;
}

.bb1c {
  width: 90%;
  height: 10%;
}

.bb1d {
  width: 100%;
  height: 70%;
}
    
```
