---
id: 5d822fd413a79914d39e98cc
title: Hatua ya 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Welcome to the CSS Variables Skyline project! Start by adding a `link` element that links your `styles.css` file within the `head` element.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Hupaswi kubadilisha tagi zako za `head` zilizopo. Hakikisha kuwa hukufuta tagi ya kufunga.

```js
const headElements = document.querySelectorAll('head');
assert.strictEqual(headElements?.length, 1);
```

You should have one void `link` element.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinafaa kuwa na sifa ya `href` yenye thamani `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
