---
id: 5d822fd413a79914d39e98cf
title: Step 4
challengeType: 0
dashedName: step-4
---

# --description--

Unaweza kuona `body` (ni kisanduku kilicho ndani kabisa kwenye ukurasa wako); kisanduku kilichozunguka ni kipengele cha `html`. Fanya `body` ijaze eneo lote la kuona kwa kumpa `height` ya `100vh`. Ondoa `margin` ya kawaida kutoka kwa `body` kwa kuweka `margin` kuwa `0`. Mwishowe, weka sifa ya `overflow` kuwa `hidden` kuficha vikanda vya kuvinjari vinavyoonekana wakati kitu kinazidi eneo la kuona.

# --hints--

Unapaswa kutumia kichaguzi cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

`body` yako inapaswa kuwa na `height` ya `100vh`.

```js
const bodyStyles = new __helpers.CSSHelp(document).getStyle('body');
assert.equal(bodyStyles?.height, '100vh');
```

`body` yako inapaswa kuwa na `margin` ya `0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.margin, '0px');
```

`body` inapaswa kuwa na sifa ya `overflow` iliyowekwa kuwa `hidden`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.overflow, 'hidden');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

--fcc-editable-region--


--fcc-editable-region--
    
```

