---
id: 5d822fd413a79914d39e98d2
title: Hatua 7
challengeType: 0
dashedName: step-7
---

# --description--

Panga ndani `div` yenye darasa la vitu la `bb1` katika kontena la majengo ya nyuma. Fungua faili lako la `styles.css`, na mpe `.bb1` darasa la vitu la `width` la `10%` na `height` la `70%`. "bb" inasimama kwa "background building", hili litakuwa jengo lako la kwanza.

# --hints--

Unapaswa kuunda kipengele kipya cha `div`.

```js
assert.equal(document.querySelectorAll('div').length, 2);
```

Unapaswa kumpa kipengele kipya cha `div` darasa la vitu la `bb1`.

```js
assert.exists(document.querySelector('div.bb1'));
```

Unapaswa kutumia kichaguzi cha darasa la vitu la `.bb1` kuipa mtindo kipengele hicho.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.bb1'));
```

Unapaswa kumpa kipengele cha `.bb1` darasa la vitu la `width` la `10%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.width, '10%');
```

Unapaswa kumpa kipengele cha `.bb1` darasa la vitu la `height` la `70%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.bb1')?.height, '70%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">    
  <head>
    <meta charset="UTF-8">
    <title>City Skyline</title>
    <link href="styles.css" rel="stylesheet" />
  </head>

  <body>
    <div class="background-buildings">
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css
* {
  border: 1px solid black;
  box-sizing: border-box;
}

body {
  height: 100vh;
  margin: 0;
  overflow: hidden;
}

.background-buildings {
  width: 100%;
  height: 100%;
}
    
```

