---
id: 655cd899f8de09431eabb40c
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 프로젝트에서는 병합 정렬 알고리즘을 구현하면서 자료 구조를 배우게 됩니다.

이 알고리즘은 분할 정복 원리를 사용해 데이터 모음을 정렬합니다. 즉, 모음을 더 작은 하위 부분으로 '분할'하고, 하위 부분을 독립적으로 정렬해 '정복'한 다음, 정렬된 하위 부분을 병합합니다.

`merge_sort`라는 함수를 만드세요. 이 함수는 숫자 목록을 정렬하는 작업을 처리합니다.

함수 본문에서 `pass` 키워드를 사용하세요.

# --hints--

`merge_sort`라는 함수를 선언해야 합니다. 함수 본문에서 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(merge_sort)
`))})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
