---
id: 6564a9fe51964c229d5b7f4c
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

정렬할 데이터 모음을 나타내는 매개변수가 필요합니다. `array` 함수에 `merge_sort`라는 매개변수를 만드세요.

# --hints--

`merge_sort` 함수는 하나의 매개변수 `array`를 받아야 합니다.

```js
({ test: () => assert(runPython(`
    import inspect
    str(inspect.signature(merge_sort)) == '(array)'    
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort():
    pass
--fcc-editable-region--
```
