---
id: 656627b47bd2d2a4afbc945d
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이전 단계에서 중간 지점을 구했습니다. 이 지점을 사용해 `array`을 두 부분으로 나누고 각 부분을 새 변수에 할당할 수 있습니다.

슬라이스 구문을 사용해 `array`의 왼쪽 절반을 추출하고 `left_part`라는 변수에 할당하세요.


# --hints--

`left_part` 함수 안에 `merge_sort`라는 변수가 있어야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*(?!=)/m);
  }
})
```

`array[:middle_point]`를 `left_part` 변수에 할당해야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /left_part\s*=\s*array\s*\[\s*:\s*middle_point\s*\]/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2

--fcc-editable-region--
```
