---
id: 656639771fed09b5c0e8fe71
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

`array` 리스트를 두 개의 별도 리스트로 나눈 후, 각 리스트를 계속 나누어 모든 요소가 각각의 리스트에 단독으로 있게 하세요. 숫자 하나만 있는 리스트는 항상 정렬된 상태입니다.

이를 위해 함수 안에서 재귀적으로 `merge_sort`을 호출하세요. 지금은 인수를 전달하지 마세요.

# --hints--

함수 본문 맨 아래에서 `merge_sort`를 호출해야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
