---
id: 65670d1ef177e7e2b76d9528
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

이전 단계에서 만든 `if`와 `else` 문장은 정렬된 배열에 요소를 할당합니다.

정렬된 배열에 할당된 각 요소는 목록에서 인덱스를 차지합니다. 따라서 요소가 없는 다음 인덱스로 이동해야 합니다.

`if`/`else` 블록 아래에 있지만 여전히 `while` 반복문 안에서 `sorted_index`를 1만큼 증가시키세요. 이 코드는 `if` 또는 `else` 문장 본문 안에 있으면 안 됩니다.

# --hints--

현재 `sorted_index` 값을 1만큼 증가시켜야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{8}sorted_index\s*(\+=\s*1|=\s*sorted_index\s*\+\s*1)/m);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
--fcc-editable-region--
```
