---
id: 656706afd65547d22bee0b68
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

`while` 반복문 안에서 `pass`을 `if`의 인덱스가 `left_part`의 인덱스보다 작은지 확인하는 `right_part` 문으로 바꾸세요.

`pass` 문 본문에 `if` 키워드를 사용하세요.

# --hints--

`pass` 키워드를 다음 조건을 가진 `if` 문으로 바꾸세요: `left_part[left_array_index] < right_part[right_array_index]`. `pass` 문 본문에 `if` 키워드를 추가하세요.

```js
({
 test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_body().is_equivalent("if left_part[left_array_index] < right_part[right_array_index]:\\n  pass")`))
 })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        pass
--fcc-editable-region--
```
