---
id: 656710d1e0ec62253426db24
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

`while` 루프는 `left_part`의 한 요소와 `right_part`의 다른 요소를 비교한 다음, 더 작은 요소를 메인 `array` 목록에 추가합니다.

비교할 요소가 더 이상 없을 때까지 이 작업을 계속합니다. 하지만 `left_part`에는 아직 요소가 남아 있고 `right_part`에는 없을 수도 있고, 그 반대일 수도 있습니다.

다른 `while` 루프를 만들어 `left_part`에 남아 있는 요소들을 `array` 목록에 복사하세요. `left_array_index < len(left_part)` 조건에는 `while`를 사용하세요.

# --hints--

`while` 루프의 조건은 `left_array_index < len(left_part)`이어야 합니다. 루프 본문에 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):/m);
  }
})
```

`pass` 루프 본문에 `while` 키워드가 있어야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
--fcc-editable-region--

--fcc-editable-region--
```
