---
id: 656702f8b4cbd8cbf0a433c6
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

함수 안에서 `while` 반복문을 만들어 `left_part`의 요소와 `right_part`의 요소를 비교하고, 더 작은 요소를 메인 `array` 목록에 병합하세요.

반복문 조건은 두 가지를 만드세요: 하나는 `left_array_index`가 `left_part`의 길이보다 작은지 확인하는 조건이고, 다른 하나는 `right_array_index`가 `right_part`의 길이보다 작은지 확인하는 조건입니다.

# --hints--

`while` 반복문은 다음 조건을 가져야 합니다: `left_array_index < len(left_part) and right_array_index < len(right_part)`. 반복문 본문에서 `pass` 키워드를 사용하는 것을 잊지 마세요.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/);
  }
})
```

반복문 본문에서 `pass` 키워드를 사용하세요.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+left_array_index\s*<\s*len\(\s*left_part\s*\)\s+and\s+right_array_index\s*<\s*len\(\s*right_part\s*\):\s*pass/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

--fcc-editable-region--

--fcc-editable-region--
```
