---
id: 6569beee367427c90c74899e
title: 29단계
challengeType: 20
dashedName: step-29
---

# --description--

이제 `merge_sort` 함수를 테스트할 시간입니다!

`pass`을 `numbers`라는 리스트로 바꾸고, 이 리스트에 `[4, 10, 6, 14, 2, 1, 8, 5]`을 할당하세요.

# --hints--

`pass`을 다음 값을 가진 `numbers` 리스트로 바꾸세요: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({ test: () =>  assert.equal(__userGlobals.get("numbers"), "[4, 10, 6, 14, 2, 1, 8, 5]") })
```

`pass` 문장 본문에 `if` 키워드가 없어야 합니다.

```js
({ test: () =>  assert(runPython(`
    _Node(_code).find_ifs()[0].find_body().is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")
`)) })
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--
if __name__ == '__main__':
    pass
--fcc-editable-region--
```
