---
id: 6569b4e0bd29d17d4b53b16c
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

이제 `while` 반복문 로직을 `right_part`에 대해 똑같이 복제할 거예요.

`while`이 `right_array_index`보다 작을 때 실행되는 `len(right_part)` 반복문을 만들고, 반복문 본문에 `pass` 키워드를 사용하세요.

# --hints--

`while` 반복문 조건은 `right_array_index < len(right_part)`이어야 합니다. 반복문 본문에 `pass` 키워드를 사용하는 것도 잊지 마세요.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /^\s{4}while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):/m);
  }
})
```

`pass` 반복문 본문에 `while` 키워드가 있어야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /while\s+right_array_index\s*<\s*len\(\s*right_part\s*\):[^}]*\bpass\b/);
  }
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--

--fcc-editable-region--
```
