---
id: 6569b743630ee592a65a7e2f
title: 26단계
challengeType: 20
dashedName: step-26
---

# --description--

`merge_sort()` 함수를 테스트하기 전에, `array`의 길이가 `1` 이하일 때 함수 실행을 정지하는 기본 사례를 만들어야 합니다.

이 기본 사례가 재귀 호출을 정지시킵니다. 이것이 없으면, 리스트가 정렬되었거나 요소가 없을 때도 병합 정렬 작업이 계속 실행됩니다.

함수 선언 바로 뒤에, 이 조건을 가진 `if` 문을 만드세요: `len(array) <= 1`. 함수 본문에서 `pass` 키워드를 사용하세요.

# --hints--

`if`인지 확인하기 위해 `len(array) <= 1` 문을 만들어야 합니다. `pass` 문 본문에서 `if` 키워드를 사용하는 것을 잊지 마세요.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_function("merge_sort").find_ifs()[0].find_conditions()[0].is_equivalent("len(array) <= 1")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):

--fcc-editable-region--
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
```
