---
id: 6569c166d708dcdf7c8fd34c
title: 31단계
challengeType: 20
dashedName: step-31
---

# --description--

`print()` 함수를 다시 호출하여 `numbers` 리스트를 출력하세요. 이렇게 하면 콘솔에 정렬되지 않은 리스트가 출력됩니다.

# --hints--

코드에 `print(numbers)`가 있어야 합니다.

```js
assert.match(code, /^\s{4}print\s*\(\s*numbers\s*\)/m)
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--
if __name__ == '__main__':
    numbers = [4, 10, 6, 14, 2, 1, 8, 5]
    print('Unsorted array:')
    
--fcc-editable-region--
```
