---
id: 6577320da0d4c2e594d418e2
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

지난 단계에서 추가한 `right_part` 호출에 `merge_sort()`을 인수로 전달하세요.

# --hints--

마지막 `right_part` 호출에 `merge_sort()`를 인수로 전달해야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*right_part\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort()
--fcc-editable-region--
```
