---
id: 6577254891048c8f2c19e961
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

`merge_sort()` 함수를 다시 호출하세요. 지금은 인수를 전달하지 마세요.

# --hints--

이전 함수 호출 후에 `merge_sort()` 함수를 다시 호출해야 합니다.

```js
({
  test: () => {
    const transformedCode = code.replace(/\r/g, "");
    const merge_sort = __helpers.python.getDef("\n" + transformedCode, "merge_sort");
    const { function_body } = merge_sort;

    assert.match(function_body, /merge_sort\s*\(\s*\)/);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)

--fcc-editable-region--
```
