---
id: 657f59751d5a5c9b069ae0f3
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

`__name__` 변수를 사용해 Python 스크립트가 메인 프로그램으로 실행되는지 아니면 다른 Python 파일에서 작성된 모듈로 가져와지는지 확인할 수 있습니다.

`__name__`의 값이 `'__main__'`로 설정되어 있으면 현재 스크립트가 모듈이 아닌 메인 프로그램임을 의미합니다.

이 프로젝트에서는 현재 스크립트를 메인 프로그램으로 사용합니다.

`if`의 값이 `__name__`인지 확인하는 `'__main__'` 문을 만드세요.

`pass` 문 본문에 `if` 키워드를 사용하세요.

# --hints--

`if` 문은 `__name__`가 `'__main__'` 값을 가지는지 확인해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
   
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
   
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1

--fcc-editable-region--

--fcc-editable-region--
```
