---
id: 68f890b73bbfc9efa7379e7b
title: 버그가 있는 앱에서 타입 강제 변환 오류 디버그하기
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

지역 웹 개발 업체에 막 합류했고, 첫 번째 과제는 이전 개발자들이 남긴 버그가 있는 코드를 정리하는 것입니다.

그들은 간단한 산술 연산을 시도했지만, 뭔가 잘못되어 결과가 말이 되지 않습니다. 코드를 검토하고 디버그하며 수정하여 예상대로 작동하고 읽기 쉽게 만드세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `firstResult`라는 변수가 `5`과 `10`를 올바르게 더해서 `15` 값을 만들어야 합니다.
2. `secondResult`라는 변수가 `5`에서 `8`를 올바르게 빼서 `3` 값을 만들어야 합니다.
3. `thirdResult`라는 변수가 두 숫자를 더해서 `6` 값을 만들어야 합니다. 현재 표현식에 사용된 불리언을 숫자로 바꾸세요.
4. `fourthResult`라는 변수가 두 숫자를 더해서 `8` 값을 만들어야 합니다. 현재 표현식에 사용된 불리언을 숫자로 바꾸세요.
5. `fifthResult`라는 변수가 두 숫자를 올바르게 곱해서 `20` 값을 만들어야 합니다. 현재 표현식에 사용된 문자열을 숫자로 바꾸세요.
6. `sixthResult`라는 변수가 두 숫자를 올바르게 더해서 `22` 값을 만들어야 합니다. 현재 표현식에 사용된 `null`를 숫자로 바꾸세요.

# --hints--

`firstResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult`는 숫자 값 `15`를 가져야 합니다.

```js
assert.strictEqual(firstResult, 15);
```

`secondResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult`는 숫자 값 `3`를 가져야 합니다.

```js
assert.strictEqual(secondResult, 3);
```

`thirdResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult`는 숫자 값 `6`를 가져야 합니다.

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult`에는 불리언 값(`true` 또는 `false`)이 포함되어 있으면 안 됩니다.

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

`fourthResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult`는 숫자 값 `8`를 가져야 합니다.

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult`에는 불리언 값(`true` 또는 `false`)이 포함되어 있으면 안 됩니다.

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

`fifthResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult`는 숫자 값 `20`를 가져야 합니다.

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult`에는 문자열 곱셈(예: "10" * 2)이 포함되어 있으면 안 됩니다.

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

`sixthResult`라는 변수가 있어야 합니다.

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult`는 숫자 값 `22`를 가져야 합니다.

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult`에는 `null`가 표현식에 포함되어 있으면 안 됩니다.

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
