---
id: 66db529d37ad966480ebb633
title: Build a Lunch Picker Program
challengeType: 25
dashedName: build-a-lunch-picker-program
---

# --description--

In this lab, you'll build a program that helps in managing lunch options. You'll work with an array of lunches, add and remove items from the array, and randomly select a lunch option.

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. You should create a `lunches` variable and assign it an empty array that will be used to store lunch items.

2. You should create a function `addLunchToEnd` that takes a string parameter and adds it to the end of the `lunches` array and returns a string in the format: `"[Lunch Item] added to the end of the lunch menu."`

3. You should create a function `addLunchToStart` that takes a string parameter and adds it to the start of the `lunches` array and returns a string in the format: `"[Lunch Item] added to the start of the lunch menu."`

4. You should create a function `removeLastLunch` that removes the last lunch item from the `lunches` array and:
   - If successful, returns a string in the format: `"[Lunch Item] removed from the end of the lunch menu."`
   - If the `lunches` array is empty, returns a string: `"No lunches to remove."`

5. You should create a function `removeFirstLunch` that removes the first lunch item from the `lunches` array and:
   - If successful, returns a string in the format: `"[Lunch Item] removed from the start of the lunch menu."`
   - If the `lunches` array is empty, returns a string: `"No lunches to remove."`

6. You should create a function `getRandomLunch` that selects a random lunch item from the `lunches` array and:
   - If successful, returns a string in the format: `"Randomly selected lunch: [Lunch Item]"`
   - If the `lunches` array is empty, returns a string: `"No lunches available."`

7. You should create a function `showLunchMenu` that:
   - If there are items in the `lunches` array, returns a string in the format: `"Menu items: [Lunch Item], [Lunch Item]...`
   - If the `lunches` array is empty, returns a string: `"The menu is empty."`

# --hints--

You should declare a variable `lunches` and assign it an empty array to store the lunch items.

```js
assert.isArray(lunches);
```

Визначте функцію `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

The function `addLunchToEnd` should take a single argument.

```js
assert.lengthOf(addLunchToEnd, 1);
```

`addLunchToEnd("Tacos")` має повернути рядок `"Tacos added to the end of the lunch menu."`.

```js
assert.strictEqual(addLunchToEnd("Tacos"), "Tacos added to the end of the lunch menu.");

// prevent hardcoding
assert.strictEqual(addLunchToEnd("Pizza"), "Pizza added to the end of the lunch menu.");
assert.strictEqual(addLunchToEnd("Burger"), "Burger added to the end of the lunch menu.");
```

Визначте функцію `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

The function `addLunchToStart` should take a single argument.

```js
assert.lengthOf(addLunchToStart, 1);
```

`addLunchToStart("Sushi")` має повернути рядок `"Sushi added to the start of the lunch menu."`.

```js
assert.strictEqual(addLunchToStart("Sushi"), "Sushi added to the start of the lunch menu.");

// prevent hardcoding
assert.strictEqual(addLunchToStart("Salad"), "Salad added to the start of the lunch menu.");
assert.strictEqual(addLunchToStart("Pasta"), "Pasta added to the start of the lunch menu.");
```

Визначте функцію `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

Якщо масив `lunches` порожній, то функція `removeLastLunch` має повернути рядок `"No lunches to remove."`.

```js
lunches = [];
assert.strictEqual(removeLastLunch(), "No lunches to remove.");
```

Якщо `lunches = ["Stew", "Soup", "Toast"]`, то функція `removeLastLunch` має повернути рядок `"Toast removed from the end of the lunch menu."`.

```js
lunches = ["Stew", "Soup", "Toast"];
assert.strictEqual(removeLastLunch(), "Toast removed from the end of the lunch menu.");

// prevent hardcoding
assert.strictEqual(removeLastLunch(), "Soup removed from the end of the lunch menu.");
assert.strictEqual(removeLastLunch(), "Stew removed from the end of the lunch menu.");
```

Якщо масив `lunches` порожній, то функція `removeFirstLunch` має повернути рядок `"No lunches to remove."`.

```js
lunches = [];
assert.strictEqual(removeFirstLunch(), "No lunches to remove.");
```

Якщо `lunches = ["Salad", "Eggs", "Cheese"]`, то функція `removeFirstLunch` має повернути рядок `"Salad removed from the start of the lunch menu."`.

```js
lunches = ["Salad", "Eggs", "Cheese"];
assert.strictEqual(removeFirstLunch(), "Salad removed from the start of the lunch menu.");

// prevent hardcoding
assert.strictEqual(removeFirstLunch(), "Eggs removed from the start of the lunch menu.");
assert.strictEqual(removeFirstLunch(), "Cheese removed from the start of the lunch menu.");
```

Визначте функцію `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

Якщо масив `lunches` порожній, то функція `getRandomLunch` має повернути рядок `"No lunches available."`.

```js
lunches = [];
assert.strictEqual(getRandomLunch(), "No lunches available.");
```

Якщо `lunches = ["Fish", "Fries", "Roast"]`, то функція `getRandomLunch` має повернути рядок у форматі `"Randomly selected lunch: [Lunch Item]"`.

```js
const temp = Math.random;

lunches = ["Fish", "Fries", "Roast"];

// check that it correctly outputs the first item
Math.random = () => 0;
assert.equal(getRandomLunch(), `Randomly selected lunch: ${lunches[0]}`);

// second item
Math.random = () => 0.4;
assert.equal(getRandomLunch(), `Randomly selected lunch: ${lunches[1]}`);

// third item
Math.random = () => 0.8;
assert.equal(getRandomLunch(), `Randomly selected lunch: ${lunches[2]}`);

// restore Math.random
Math.random = temp;
```

Визначте функцію `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

Якщо масив `lunches` порожній, то функція `showLunchMenu` має повернути рядок `"The menu is empty."`.

```js
lunches = [];
assert.strictEqual(showLunchMenu(), "The menu is empty.");
```

Якщо `lunches = ["Greens", "Corns", "Beans"]`, то функція `showLunchMenu` має повернути рядок у форматі `"Menu items: Greens, Corns, Beans"`.

```js
lunches = ["Greens", "Corns", "Beans"];
assert.strictEqual(showLunchMenu(), "Menu items: Greens, Corns, Beans");

lunches = ["Pizza", "Burger", "Fries"];
assert.strictEqual(showLunchMenu(), "Menu items: Pizza, Burger, Fries");
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the list (at the end)
function addLunchToEnd(newLunch) {
  lunches.push(newLunch);
  return `${newLunch} added to the end of the lunch menu.`;
}

// Add a new lunch to the start of the list
function addLunchToStart(newLunch) {
  lunches.unshift(newLunch);
  return `${newLunch} added to the start of the lunch menu.`;
}

// Remove the last lunch from the list
function removeLastLunch() {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    return `${removedLunch} removed from the end of the lunch menu.`;
  } else {
    return "No lunches to remove.";
  }
}

// Remove the first lunch from the list
function removeFirstLunch() {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    return `${removedLunch} removed from the start of the lunch menu.`;
  } else {
    return "No lunches to remove.";
  }
}

// Function to get a random lunch from the list
function getRandomLunch() {
  if (lunches.length === 0) {
    return "No lunches available.";
  }
  const randomIndex = Math.floor(Math.random() * lunches.length);
  const randomLunch = lunches[randomIndex];
  return `Randomly selected lunch: ${randomLunch}`;
}

// New function to display the lunches as a numbered menu using a for loop
function showLunchMenu() {
  if (lunches.length === 0) {
    return "The menu is empty.";
  }
  return `Menu items: ${lunches.join(', ')}`;
}

addLunchToEnd("Tacos");
addLunchToStart("Sushi");
removeLastLunch();
removeFirstLunch();
getRandomLunch();
showLunchMenu();

```
