---
id: 6852f01116912861c7d1b0e6
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

在 `__init__` 方法内，你需要将参数值赋值给对象的属性，以便每个电子邮件可以保存其信息。

在你的 `__init__` 方法中，删除 `pass` 关键字，并将参数赋值给 `self.sender`、`self.receiver`、`self.subject` 和 `self.body`。

# --hints--

你应该将 `sender` 参数赋值给 `self.sender`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.sender").is_equivalent("self.sender = sender")`))
    }

})
```

你应该将 `receiver` 参数赋值给 `self.receiver`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.receiver").is_equivalent("self.receiver = receiver")`))
    }

})
```

你应该将 `subject` 参数赋值给 `self.subject`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.subject").is_equivalent("self.subject = subject")`))
    }

})
```

你应该将 `body` 参数赋值给 `self.body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").find_variable("self.body").is_equivalent("self.body = body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        pass
--fcc-editable-region--
```
