---
id: 6852f0c3fcedb964ab959921
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

现在让我们测试 `read` 属性是否已正确添加到你的 `Email` 类。由于你已经有了前面步骤中的 email 对象，打印 `read` 属性以查看它现在的默认值是否为 `False`。

# --hints--

你应该打印 `email_obj.read` 来测试读取属性。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(email_obj.read)")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)

--fcc-editable-region--
```
