---
id: 6852f11d412c29665985a915
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

既然用户可以发送电子邮件，他们需要一种方法来保存收到的电子邮件。每个用户都应该有一个收件箱来收集他们的电子邮件。

为你的 `User` 类的 `__init__` 方法添加一个 `inbox` 属性，并将其初始化为空列表 `[]`。这将保存用户接收的所有电子邮件。

# --hints--

你应该为你的 `User` 类的 `__init__` 方法添加 `self.inbox = []`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
```
