---
id: 6852ef95ccc27660275692cd
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你的 `Email` 类需要在创建时保存每封邮件的相关信息。

在你的 `Email` 类中，删除现有的 `pass` 关键字并创建接受 `self`、`sender`、`receiver`、`subject` 和 `body` 作为参数的 `__init__` 方法。

# --hints--

你应该在 `Email` 类中定义一个 `__init__` 方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__init__")`))
    }

})
```

你的 `__init__` 方法应有参数：`self`、`sender`、`receiver`、`subject` 和 `body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("__init__").has_args("self, sender, receiver, subject, body")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    pass
--fcc-editable-region--
```
