---
id: 6852f06fdde1b4634a263f03
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

现在让我们通过创建一个 `Email` 对象并查看它是否正确保存信息来测试 `Email` 类。你将打印几个属性以验证一切是否正常。

创建一个名为 `email_obj` 的电子邮件对象，发件人为 `alice@example.com`，收件人为 `bob@example.com`，主题为 `Hello`，正文为 `Hi Bob!`。然后分别打印 `sender` 和 `subject` 属性作为独立的打印语句，以验证它们是否正确保存。

# --hints--

你应该创建一个名为 `email_obj` 的电子邮件对象，并带有指定的参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_variable("email_obj").is_equivalent("email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')")`))
    }

})
```

你应该打印 `email_obj.sender`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.sender)")`))
    }

})
```

你应该打印 `email_obj.subject`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("print(email_obj.subject)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--

--fcc-editable-region--
```
