---
id: 6852ea0c3cadb54e897acdac
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在本次研讨会中，你将构建一个*Email Simulator*，用于模拟不同用户之间发送、接收和管理电子邮件。你将学习关于类、对象以及如何以面向对象的方式组织代码。

首先使用 `class` 关键字创建一个名为 `Email` 的类。

# --hints--

你应该创建一个名为 `Email` 的类。

```js
({
    test: () => 
    {
        assert(runPython(`(_Node(_code).has_class("Email"))`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
