---
id: 6853f2da9b8ed891e9fdc5fc
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

现在让我们显示发件人和收件人简介。要显示谁发送和接收了电子邮件，请为 `display_full_email` 方法添加两个打印语句，格式如下：

- `From: sender` 其中 `sender` 被替换为发送者的名称
- `To: receiver` 其中 `receiver` 被替换为接收者的名称

# --hints--

你应该使用 f-strings 来打印发送者的名字，格式为 `From: {self.sender.name}`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'From: {self.sender.name}')")`));
    }
})
```

你应该使用 f-strings 来打印接收者的名字，格式为 `To: {self.receiver.name}`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("print(f'To: {self.receiver.name}')")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')

--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
