---
id: 6853f90fd03acc9fd75ff860
title: 步骤46
challengeType: 20
dashedName: step-46
---

# --description--

用户应该能够通过他们的 User 对象查看收件箱、读取电子邮件并删除电子邮件。

为查看收件箱，为 `User` 类添加一个名为 `check_inbox` 的方法。该方法应打印带有用户姓名的个性化头部，然后显示他们的所有邮件。

头部应格式化为：`\nName's Inbox:`，其中 `Name` 替换为用户的名字。

打印头部后，调用用户收件箱上的 `list_emails()` 方法。

# --hints--

你应该在 `User` 类中定义一个名为 `check_inbox` 的方法，该方法以 `self` 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").has_args("self")`))
    }

})
```

在 `check_inbox` 中，你应该打印一个包含用户姓名的头部，格式为 `\n[user]'s Inbox:`。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[0].is_equivalent('print(f"\\\\n{self.name}\\\'s Inbox:")')`))
    }
})
```

在 `check_inbox` 中，你应该调用 `self.inbox.list_emails()` 来显示邮件。

```js
({ 
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("check_inbox").find_body()[1].is_equivalent("self.inbox.list_emails()")`))
    }
})

```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

--fcc-editable-region--
    
--fcc-editable-region--

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

```
