---
id: 6853f34ff93594934f6636d9
title: 步骤 27
challengeType: 20
dashedName: step-27
---

# --description--

让我们为你的 `Email` 类添加一个字串表现，以便我们可以显示简短的电子邮件摘要。

为你的 `Email` 类添加一个以 `self` 作为参数的 `__str__` 方法。

# --hints--

你应该在 `Email` 类中定义一个 `__str__` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("__str__")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
