---
id: 68540a534c56aeabae6afd58
title: 步骤 50
challengeType: 20
dashedName: step-50
---

# --description--

现在你将模拟用户之间发送一些电子邮件。

在你的 `main` 函数中，创建用户后，使用 `send_email` 方法来：

- 让 Tory 向 `ramy` 用户对象发送主题为 `Hello` 且正文为 `Hi Ramy, just saying hello!` 的电子邮件。
- 让 Ramy 发送一封电子邮件给 `tory` 用户对象，主题为 `Re: Hello`，正文为 `Hi Tory, hope you are fine.`

# --hints--

你应该让 Tory 给 `ramy` 发送一封电子邮件，主题为 `Hello`，内容为 `Hi Ramy, just saying hello!`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body()[2].is_equivalent("tory.send_email(ramy, 'Hello', 'Hi Ramy, just saying hello!')")`))
    }

})
```

你应该让 Ramy 给 `tory` 发送一封电子邮件，主题为 `Re: Hello`，正文为 `Hi Tory, hope you are fine.`

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_function("main").find_body()[3].is_equivalent("ramy.send_email(tory, 'Re: Hello', 'Hi Tory, hope you are fine.')")`))
    }

})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.timestamp = datetime.datetime.now()
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f"Received: {self.timestamp.strftime('%Y-%m-%d %H:%M')}")
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject} | Time: {self.timestamp.strftime('%Y-%m-%d %H:%M')}"
class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')


    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)
        print(f'Email sent from {self.name} to {receiver.name}!\n')

    def check_inbox(self):
        print(f"\n{self.name}'s Inbox:")
        self.inbox.list_emails()

    def read_email(self, index):
        self.inbox.read_email(index)

    def delete_email(self, index):
        self.inbox.delete_email(index)

--fcc-editable-region--
def main():
    tory = User('Tory')
    ramy = User('Ramy')

--fcc-editable-region--
if __name__ == '__main__':
    main()
```
