---
id: 6859400f8f7f872e50e1238c
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

`User` 类需要一个方法来向其他用户发送电子邮件。发送电子邮件时，你将创建一个新的 `Email` 对象并将其添加到接收者的收件箱。

在你的 `User` 类中创建一个名为 `send_email` 的方法。该方法应接受 `receiver`、`subject` 和 `body` 的参数。 

在该方法内，使用以下参数值创建一个新的 `Email` 对象，并将其赋值给名为 `email` 的变量：

- `sender`：`self`（当前用户）
- `receiver`：`receiver`（接收电子邮件的用户）
- `subject`：`subject`（电子邮件的主题）
- `body`：`body`（电子邮件的体）

# --hints--

你应该在 `User` 类中定义一个 `send_email` 方法。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").has_function("send_email")`))
    }

})
```

`send_email` 方法应具有参数：`self`、`receiver`、`subject` 和 `body`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("send_email").has_args("self, receiver, subject, body")`))
    }

})
```

该方法应使用正确的参数创建一个名为 `email` 的 `Email` 对象。

```js
({
  test: () => {
    assert(runPython(`
cond = _Node(_code).find_class("User").find_function("send_email").find_body()[0]
cond.is_equivalent("email = Email(sender=self, receiver=receiver, subject=subject, body=body)") or cond.is_equivalent("email = Email(self, receiver, subject, body)")
    `));
  }
});
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name

--fcc-editable-region--
```
