---
id: 68593ea22eac7e228ed585ce
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

电子邮件应跟踪是否已被 `read`。为 `__init__` 方法添加一个 `read` 属性，并将其默认设置为 `False`，因为新邮件默认是未读的。

# --hints--

你应该在 `__init__` 方法中将 `False` 赋值给 `self.read`。

```js
({ 
    test: () => runPython(`
    assert _Node(_code).find_class("Email").find_function("__init__").find_variable("self.read").is_equivalent("self.read = False")
    `) 
})
```


# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
```
