---
id: 685c1eba966c6396850736ea
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

虽然用户可以发送电子邮件并拥有收件箱，但我们需要一个专用的类来高效管理收件箱操作。 

`Inbox` 类将保存一个电子邮件列表，并提供添加新电子邮件、列出所有电子邮件和管理它们的方法。

创建一个名为 `Inbox` 的新建类，该类包含一个 `__init__` 方法，用于初始化一个名为 `emails` 的空列表。

# --hints--

你应该创建一个 `Inbox` 类。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_class("Inbox")`));
    }
})
```

`Inbox` 类应有一个接受 `self` 作为参数的 `__init__` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").has_args("self")`));
    }
})
```

你应该在 `Inbox` 类的 `__init__` 方法中初始化 `self.emails = []`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("__init__").find_variable("self.emails").is_equivalent("self.emails = []")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--

--fcc-editable-region--
```
