---
id: 685c1e32254e5993a5c8823e
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

既然你已经完成了 `Email` 类的初始设置，现在是时候创建能够发送和接收电子邮件的用户了。每个用户都会有一个名字，并且能够执行电子邮件操作。

创建一个带有 `__init__` 方法的 `User` 类，该方法接受 `self` 和 `name` 作为参数。在类中，将 `name` 参数赋值给 `self.name`。

# --hints--

你应该创建一个 `User` 类。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_class("User")`))
    }

})
```

`User` 类应有一个 `__init__` 方法，该方法接受 `self` 和 `name` 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").has_args("self, name")`))
    }

})
```

你应该将 `name` 参数赋值给 `self.name`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.name").is_equivalent("self.name = name")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--

--fcc-editable-region--
```
