---
id: 685d2fd1f309b9191f257c55
title: 步骤 40
challengeType: 20
dashedName: step-40
---

# --description--

在将时间戳集成到我们的电子邮件系统之前，让我们先练习处理日期时间格式。`datetime.datetime.now()` 函数提供当前的日期和时间，我们可以使用 `strftime()` 方法以不同的方式格式化它。

以下是 `strftime()` 如何与格式代码一起工作的方式：

```python
now = datetime.datetime.now()
print(now.strftime("%Y-%m-%d"))  # Output: 2024-03-15 (year-month-day with - separator)
```

格式代码如 `%Y`（年）、`%m`（月）、`%d`（日）告诉 `strftime()` 要包含什么，你可以在它们之间添加像 `-` 这样的分隔符。

在你的代码底部，创建一个名为 `current_time` 的变量并赋值为 `datetime.datetime.now()`。然后使用 `strftime()` 以小时:分钟:秒的格式打印时间，使用 `:` 作为分隔符。 

使用这些格式代码：`%H` 表示小时（24 小时制），`%M` 表示分钟，`%S` 表示秒。

# --hints--

你应该创建一个名为 `current_time` 的变量并赋值为 `datetime.datetime.now()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_stmt("current_time = datetime.datetime.now()")`));
    }
})
```

你应该使用 `strftime()` 以 `%H:%M:%S` 形式打印时间。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call("print(current_time.strftime('%H:%M:%S'))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
import datetime

class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        self.emails[actual_index].display_full_email()

    def delete_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        actual_index = index - 1
        if actual_index < 0 or actual_index >= len(self.emails):
            print('Invalid email number.\n')
            return
        del self.emails[actual_index]
        print('Email deleted.\n')

--fcc-editable-region--

--fcc-editable-region--
```
