---
id: 6887550fe36f64e27074475a
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

让 Alice 给 Bob 发送电子邮件，以确认电子邮件发送是否正常。

使用 Alice 的 `send_email` 方法向 Bob 发送主题为 `"Hello"` 且正文为 `"Hi Bob, how are you?"` 的电子邮件。

# --hints--

你应该使用正确的参数调用 `alice.send_email()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).has_call('alice.send_email(bob, "Hello", "Hi Bob, how are you?")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

--fcc-editable-region--
```
