---
id: 688754d6d59578df1313cb3c
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

现在是连接 `User` 和 `Inbox` 类的时候了，这样邮件才能真正发送！ 

用户需要拥有正确的 Inbox 对象，而不是简单的列表，且 `send_email` 方法应将邮件发送到接收者的收件箱。

更新 `User` 类的 `__init__`，使每个用户获得一个实际的 Inbox 对象，而不是之前创建的空列表。 

# --hints--

你应该更新 User 类以初始化 `self.inbox = Inbox()`。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("User").find_function("__init__").find_variable("self.inbox").is_equivalent("self.inbox = Inbox()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = []
--fcc-editable-region--

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

class Inbox:
    def __init__(self):
        self.emails = []

```
