---
id: 688798fccea91f522d1f245e
title: 步骤 33
challengeType: 20
dashedName: step-33
---

# --description--

当收件箱不为空时，你将尝试访问给定索引处的电子邮件并显示它。 

请记住，在 `list_emails` 方法中，你显示的电子邮件编号是从 1 开始的，但 Python 中的列表索引是从 0 开始的。因此，你需要通过减去 1 将基于 1 的索引转换为基于 0 的索引。

在 `read_email` 方法中，从 `index` 参数中减去 1 并将其保存到名为 `actual_index` 的变量中。 


# --hints--

你应该从 `index` 参数中减去 1 并将其保存到名为 `actual_index` 的变量中。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("read_email").find_body().find_variable("actual_index").is_equivalent("actual_index = index - 1")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

    def display_full_email(self):
        self.mark_as_read()
        print('\n--- Email ---')
        print(f'From: {self.sender.name}')
        print(f'To: {self.receiver.name}')
        print(f'Subject: {self.subject}')
        print(f'Body: {self.body}')
        print('------------\n')

    def __str__(self):
        status = 'Read' if self.read else 'Unread'
        return f"[{status}] From: {self.sender.name} | Subject: {self.subject}"
class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

    def list_emails(self):
        if not self.emails:
            print('Your inbox is empty.\n')
            return
        print('\nYour Emails:')
        for i, email in enumerate(self.emails, start=1):
            print(f'{i}. {email}')

--fcc-editable-region--
    def read_email(self, index):
        if not self.emails:
            print('Inbox is empty.\n')
            return
        
--fcc-editable-region--

```
