---
id: 68875554c6c32de788afc64e
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

用户需要一种正确 `读取` 他们电子邮件的方法。

为你的 `Email` 类添加一个只接受 `self` 作为参数的 `display_full_email` 方法。在此方法内部，调用 `mark_as_read()` 方法以在有人查看邮件时将其标记为已读。

此方法用于以美观的格式显示电子邮件的完整内容。

# --hints--

你应该在 `Email` 类中定义一个 `display_full_email` 方法。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").has_function("display_full_email")`));
    }
})
```

`display_full_email` 方法应仅接受 `self` 作为参数。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_args("self")`));
    }
})
```

该方法应调用 `self.mark_as_read()` 来将电子邮件标记为已读取。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_class("Email").find_function("display_full_email").has_call("self.mark_as_read()")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True
--fcc-editable-region--
    
--fcc-editable-region--

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

```
