---
id: 688754f3494b17e0b352d67d
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

现在是测试我们完整电子邮件系统的时候了！让我们创建一些用户，看看电子邮件功能性如何运行。

创建两个用户对象：`alice`，名称为 `"Alice"`，和 `bob`，名称为 `"Bob"`。这将演示用户如何被创建并相互交互。

# --hints--

你应该创建一个名为 `alice` 且名称为 `"Alice"` 的用户对象。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("alice").is_equivalent('alice = User("Alice")')`));
    }
})
```

你应该创建一个名为 `bob` 且名字为 `"Bob"` 的用户对象。

```js
({
    test: () => {
        assert(runPython(`_Node(_code).find_variable("bob").is_equivalent('bob = User("Bob")')`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []

    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--

--fcc-editable-region--
```
