---
id: 68b9366fae8b13c1278761b1
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

你的收件箱需要一种接收新邮件的方式。当有人向用户发送邮件时，邮件应被添加到他们的收件箱。

为你的 `Inbox` 类添加一个名为 `receive_email` 的方法，该方法以 `self` 和一个电子邮件对象 `email` 作为参数。在方法体内，使用 `append` 方法将 `email` 追加到 `emails` 列表中。

# --hints--

你应该在 `Inbox` 类中定义一个名为 `receive_email` 的方法，该方法接受 `self` 和 `email` 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_args("self, email")`))
    }
})
```

你应该将 `email` 追加到 `self.emails`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Inbox").find_function("receive_email").has_stmt("self.emails.append(email)")`))
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)

--fcc-editable-region--
class Inbox:
    def __init__(self):
        self.emails = []

    
--fcc-editable-region--
```
