---
id: 68b9350183ff05bd106c79e4
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在，测试你在 `Email` 类中创建的 `mark_as_read` 方法。使用该方法对现有的 `Email` 对象 `email_obj` 进行操作以更改状态。标记邮件为已读后，打印 `email_obj` 的 `read` 属性以确认更改。

# --hints--

你应该在 `email_obj` 上调用 `mark_as_read` 方法以将邮件标记为已读取。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).has_call("email_obj.mark_as_read()")`))
    }

})
```

你应该在调用 `mark_as_read` 方法后打印 `email_obj` 的 `read` 属性。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).is_ordered("email_obj.mark_as_read()", "print(email_obj.read)")`))
    }

})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)

--fcc-editable-region--
```
