---
id: 68b9a011daf902ff1eba40a8
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

既然电子邮件已被接收，移除对象创建、方法调用和打印语句。

# --hints--

你的代码中不应有 `alice = User("Alice")`。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("alice = User(\\"Alice\\")")`));
    }
})
```

你的代码中不应有 `bob = User("Bob")`。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_stmt("bob = User(\\"Bob\\")")`));
    }
})
```

你的代码中不应包含 `alice.send_email(bob, "Hello", "Hi Bob, how are you?")`。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("alice.send_email(bob, \\"Hello\\", \\"Hi Bob, how are you?\\")")`));
    }
})
```

你的代码中不应有 `print(len(bob.inbox.emails))`。

```js
({
    test: () => {
        assert(!runPython(`_Node(_code).has_call("print(len(bob.inbox.emails))")`));
    }
})
```

# --seed--

## --seed-contents--

```py
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

    def mark_as_read(self):
        self.read = True

class User:
    def __init__(self, name):
        self.name = name
        self.inbox = Inbox()

    def send_email(self, receiver, subject, body):
        email = Email(sender=self, receiver=receiver, subject=subject, body=body)
        receiver.inbox.receive_email(email)

class Inbox:
    def __init__(self):
        self.emails = []
    
    def receive_email(self, email):
        self.emails.append(email)

--fcc-editable-region--
alice = User("Alice")
bob = User("Bob")

alice.send_email(bob, "Hello", "Hi Bob, how are you?")

print(len(bob.inbox.emails))
--fcc-editable-region--
```
