---
id: 68b930ee0e7b8fb81895e59a
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在你将添加一个方法来标记电子邮件为已读。当有人打开电子邮件时，你想将其状态从未读更改为已读。此方法很简单——只需将 `read` 属性设置为 `True`。

为你的 `Email` 类添加一个名为 `mark_as_read` 的方法。该方法应仅接受 `self` 作为参数，并将 `read` 属性设置为 `True`。

# --hints--

你应该在 `Email` 类中定义一个名为 `mark_as_read` 的方法，该方法只接受 `self` 作为参数。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").has_args("self")`))
    }

})
```

该方法应将 `self.read` 设置为 `True`。

```js
({
    test: () => 
    {
        assert(runPython(`_Node(_code).find_class("Email").find_function("mark_as_read").find_variable("self.read").is_equivalent("self.read = True")`))
    }

})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Email:
    def __init__(self, sender, receiver, subject, body):
        self.sender = sender
        self.receiver = receiver
        self.subject = subject
        self.body = body
        self.read = False

--fcc-editable-region--
email_obj = Email('alice@example.com', 'bob@example.com', 'Hello', 'Hi Bob!')
print(email_obj.sender)
print(email_obj.subject)
print(email_obj.read)
```
