---
id: 587d7fb7367417b2b2512c0c
title: Usar model.findOne() para retornar um único documento correspondente a partir do seu banco de dados
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` se comporta como `Model.find()`, mas retorna apenas um documento (não um array), mesmo que existam vários itens. É especialmente útil ao pesquisar por propriedades que você declarou como exclusivas.

# --instructions--

Modifique a função `findOneByFood` para encontrar apenas uma pessoa que tenha uma certa comida nos favoritos da pessoa, usando `Model.findOne() -> Person`. Use o argumento `food` da função como chave de pesquisa.

# --hints--

Você deve ter sucesso em encontrar um item

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

