---
id: 66c644c08ebcd2ef90c750bd
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

У попередніх уроках ви навчились додавати елементи в кінець масиву за допомогою методу `push` ось так:

```js
fruits.push('pear');
```

Використовуючи метод `push`, додайте рядок `"Apples"` до масиву `shoppingList`.

# --hints--

Ваш масив `shoppingList` повинен містити один елемент.

```js
assert.lengthOf(shoppingList, 1);
```

Ваш масив `shoppingList` повинен містити рядок `"Apples"`. Перевірте правильність написання та використання великих літер.

```js
assert.includeMembers(shoppingList, ['Apples']);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

--fcc-editable-region--

--fcc-editable-region--
```
