---
id: 66c63bfa4685e4e3843fa0dc
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви продовжите вивчати масиви, створюючи список покупок продуктів.

Почніть із додавання `console.log`, який виводить рядок `"Grocery shopping list"` у консоль.

# --hints--

У вашому коді має бути `console.log`.

```js
assert.match(code, /console\.log(.*)/);
```

Ваш `console.log` має виводити повідомлення `"Grocery shopping list"`. Перевірте орфографію та правильність використання великих літер.

```js
assert.match(code, /console\.log\((["'])(Grocery\s+shopping\s+list)\1\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
