---
id: 66c645b345a39bf04864dc50
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Щоб побачити оновлений список покупок, потрібно вивести в консоль поточний список покупок і коротке повідомлення.

Оскільки це повідомлення буде використовуватися багаторазово протягом практичного заняття, найкраще створити багаторазову функцію.

Створіть функцію з назвою `getShoppingListMsg`, яка приймає масив як параметр і повертає рядок `"Current Shopping List: "`, за яким іде вміст переданого масиву.

Ви можете використовувати шаблонні літерали або конкатенацію рядків з оператором `+`, щоб поєднати рядок із масивом, переданим у функцію.

# --hints--

У вашому коді має бути функція з назвою `getShoppingListMsg`.

```js
assert.isFunction(getShoppingListMsg);
```

Ваша функція `getShoppingListMsg` має повертати рядок.

```js
assert.isString(getShoppingListMsg(shoppingList));
```

Ваша функція `getShoppingListMsg` має приймати масив і повертати повідомлення `"Current Shopping List: "`, за яким іде вміст цього масиву. Перевірте правильність пробілів і орфографії.

```js
assert.oneOf(getShoppingListMsg(["Anchovies", "Olives"]), ["Current Shopping List: Anchovies,Olives", "Current Shopping List: Anchovies, Olives"]);
assert.strictEqual(getShoppingListMsg(["Apples"]), "Current Shopping List: Apples");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

--fcc-editable-region--

--fcc-editable-region--
```
