---
id: 66c64210413532ee9d3bd342
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У цій частині практичного заняття ви будете практикувати додавання продуктів до списку покупок.

Почніть із використання `console.log` для виведення повідомлення `"It will be nice to have some fruit to eat."`.

# --hints--

Ви маєте вивести повідомлення `"It will be nice to have some fruit to eat."` у консоль.

```js
assert.match(code, /console\.log\s*\(\s*(['"])(It\s+will\s+be\s+nice\s+to\s+have\s+some\s+fruit\s+to\s+eat\.)\1\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

--fcc-editable-region--

--fcc-editable-region--
```
