---
id: 66c6491fe4c8e0f16845425f
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер настав час додати ще один фрукт до списку.

Використовуючи той самий метод масиву, що й раніше, додайте рядок `"Grapes"` у кінець масиву `shoppingList`.

Потім додайте `console.log` і викличте функцію `getShoppingListMsg` з масивом `shoppingList` як аргументом всередині `console.log`, щоб побачити оновлений список у консолі.

# --hints--

Ваш масив `shoppingList` тепер має містити два елементи.

```js
assert.lengthOf(shoppingList, 2);
```

Ваш масив `shoppingList` має містити рядок `"Grapes"`.

```js
assert.equal(shoppingList[1], "Grapes");
```

Ви маєте викликати функцію `getShoppingListMsg(shoppingList)` всередині `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 2);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
