---
id: 66c64095f61166ede6e39a84
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Для цього списку покупок ви використаєте масив, щоб представити елементи, які потрібно купити.

У попередніх уроках ви навчились створювати масиви так:

```js
const fruits = ['apple', 'banana', 'orange'];
```

На цьому кроці створіть змінну з назвою `shoppingList` і призначте їй порожній масив.

# --hints--

У вас має бути змінна з назвою `shoppingList`.

```js
assert.isNotNull(shoppingList);
```

Ви маєте призначити масив змінній `shoppingList`.

```js
assert.isArray(shoppingList);
```

Ваш масив має бути порожнім.

```js
assert.isEmpty(shoppingList);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

--fcc-editable-region--

--fcc-editable-region--
```
