---
id: 66c72f4d0528bd268a82107b
title: Крок 12
challengeType: 1
dashedName: step-12
---

# --description--

Тепер настав час вивести оновлений масив `shoppingList` у консоль.

Додайте ще один `console.log` і викличте функцію `getShoppingListMsg` з масивом `shoppingList` як аргументом всередині `console.log`, щоб побачити оновлений список у консолі.

# --hints--

Ви повинні викликати `getShoppingListMsg(shoppingList)` всередині `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 4);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");

--fcc-editable-region--

--fcc-editable-region--
```
