---
id: 66c730ee6ae076281721d0b9
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

У попередніх уроках ви навчились видаляти елементи з кінця масиву за допомогою методу `pop`.

Ось нагадування, як використовувати метод `pop`:

```js
let array = [1, 2, 3, 4, 5];
array.pop();

// [1, 2, 3, 4]
console.log(array); 
```

Використайте метод `pop`, щоб видалити останній елемент з масиву `shoppingList`.

# --hints--

Ваш масив `shoppingList` тепер має містити `5` елементів.

```js
assert.lengthOf(shoppingList, 5);
```

У вашому масиві `shoppingList` більше не має бути рядка `"Potato Chips"`.

```js
assert.notInclude(shoppingList, "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

--fcc-editable-region--

--fcc-editable-region--
```
