---
id: 66c73a7798f6f62b2ae58f22
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Тепер час додати більше елементів на початок списку покупок.

Почніть із додавання інструкції `console.log`, яка виводить повідомлення `"It might be nice to get a dessert."`.

Під цією інструкцією `console` використайте правильний метод масиву, щоб додати рядок `"Chocolate Cake"` на початок масиву `shoppingList`.

Нарешті, додайте `console.log` і викличте функцію `getShoppingListMsg` з масивом `shoppingList` як аргументом всередині `console.log`, щоб побачити оновлений список у консолі.

# --hints--

Ви повинні мати інструкцію `console.log`, яка виводить повідомлення `"It might be nice to get a dessert."`.

```js
assert.match(code, /console\.log\((['"])(It\s+might\s+be\s+nice\s+to\s+get\s+a\s+dessert\.)\1\)/);
```

Ваш масив `shoppingList` має містити загалом `6` елементів.

```js
assert.lengthOf(shoppingList, 6);
```

Ви повинні використати метод `unshift`, щоб додати рядок `"Chocolate Cake"` на початок масиву `shoppingList`.

```js
assert.strictEqual(shoppingList[0], "Chocolate Cake");
```

Ви повинні викликати `getShoppingListMsg(shoppingList)` всередині `console.log`.

```js
assert.lengthOf(code.match(/console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/g), 6);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Popcorn", "Beef Jerky", "Potato Chips");
console.log(getShoppingListMsg(shoppingList));

console.log("This looks like too much junk food.");

shoppingList.pop();
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
