---
id: 66c72b0ffbc5522525768558
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

У попередніх кроках ви повторювали, як додати елемент у кінець масиву за допомогою методу `push`.

Але метод `push` приймає кілька аргументів, тож ви можете додати кілька елементів у кінець масиву ось так:

```js
array.push(item1, item2, item3);
```

У цьому кроці використайте метод `push`, щоб додати рядки `"Popcorn", "Beef Jerky", "Potato Chips"` до масиву `shoppingList`.

Порядок важливий, тому переконайтеся, що додаєте елементи у тому порядку, в якому вони наведені.

# --hints--

Ваш масив `shoppingList` має містити загалом `6` рядків.

```js
assert.lengthOf(shoppingList, 6);
```

Ваш `shoppingList` має містити рядок `"Popcorn"`.

```js
assert.strictEqual(shoppingList[3], "Popcorn");
```

Ваш `shoppingList` має містити рядок `"Beef Jerky"`.

```js
assert.strictEqual(shoppingList[4], "Beef Jerky");
```

Ваш `shoppingList` має містити рядок `"Potato Chips"`.

```js
assert.strictEqual(shoppingList[5], "Potato Chips");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

shoppingList.unshift("Vegetable Oil");
console.log(getShoppingListMsg(shoppingList));

--fcc-editable-region--

--fcc-editable-region--
```
