---
id: 66c726c34ecf1e238aa9d7d5
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

У попередніх уроках ви навчились додавати елементи на початок масиву за допомогою методу `unshift()`.

Ось нагадування, як використовувати метод `unshift()`:

```js
array.unshift(item1, item2, ..., itemX);
```

Використайте метод `unshift()`, щоб додати рядок `"Vegetable Oil"` на початок масиву `shoppingList`.

# --hints--

У вашому масиві `shoppingList` має бути три елементи.

```js
assert.lengthOf(shoppingList, 3);
```

Ви маєте використати метод `unshift()`, щоб додати `"Vegetable Oil"` на початок масиву `shoppingList`.

```js
assert.equal(shoppingList[0], "Vegetable Oil");
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");
function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

console.log(getShoppingListMsg(shoppingList));

shoppingList.push("Grapes");
console.log(getShoppingListMsg(shoppingList));

console.log("It looks like we need to get some cooking oil.");

--fcc-editable-region--

--fcc-editable-region--
```
