---
id: 66cbe2319d3845545a293a0b
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер настав час побачити повідомлення, яке виводиться в консоль.

Додайте `console.log` і викличте функцію `getShoppingListMsg` з масивом `shoppingList` як аргументом всередині `console.log`, щоб побачити повідомлення в консолі.

# --hints--

Ви повинні викликати функцію `getShoppingListMsg(shoppingList)` всередині `console.log`.

```js
assert.match(code, /console\.log\(\s*getShoppingListMsg\(\s*shoppingList\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Grocery shopping list");

const shoppingList = [];

console.log("It will be nice to have some fruit to eat.");

shoppingList.push("Apples");

function getShoppingListMsg(arr) {
  return `Current Shopping List: ${arr}`;
}

--fcc-editable-region--

--fcc-editable-region--
```
