---
id: 587d78ae367417b2b2512afd
title: Використання властивості flex-basis для встановлення ініціального розміру елемента
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

Властивість `flex-basis` задає ініціальний розмір елемента до того, як CSS зробить правки за допомогою `flex-shrink` або `flex-grow`.

Одиниці, які використовує властивість `flex-basis`, такі ж, як і в інших властивостях розміру (`px`, `em`, `%` тощо). Значення `auto` задає розмір елементів залежно від вмісту.

# --instructions--

Встановіть ініціальний розмір блоків, використовуючи `flex-basis`. Додайте CSS-властивість `flex-basis` до обох `#box-1` і `#box-2`. Задайте `#box-1` значення `10em`, а `#box-2` — значення `20em`.

# --hints--

Елемент `#box-1` має мати властивість `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Елемент `#box-1` має мати значення `flex-basis` рівне `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

Елемент `#box-2` має мати властивість `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Елемент `#box-2` має мати значення `flex-basis` рівне `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
