---
id: 587d78ad367417b2b2512afb
title: Використання властивості flex-shrink для зменшення елементів
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

До цього моменту всі властивості в завданнях застосовувалися до гнучкого контейнера (батьківського елемента гнучких елементів). Однак існує кілька корисних властивостей для гнучких елементів.

Перша — це властивість `flex-shrink`. Коли її використовують, вона дозволяє елементу зменшуватися, якщо гнучкий контейнер занадто малий. Елементи зменшуються, коли ширина батьківського контейнера менша за сумарну ширину всіх гнучких елементів у ньому.

Властивість `flex-shrink` приймає числа як значення. Чим більше число, тим більше елемент зменшиться порівняно з іншими елементами в контейнері. Наприклад, якщо один елемент має значення `flex-shrink` рівне `1`, а інший — `flex-shrink` рівне `3`, то той, у кого значення `3`, зменшиться втричі більше за іншого.

# --instructions--

Додайте CSS-властивість `flex-shrink` до обох `#box-1` і `#box-2`. Дайте `#box-1` значення `1`, а `#box-2` — значення `2`.

# --hints--

Елемент `#box-1` повинен мати властивість `flex-shrink` зі значенням `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

Елемент `#box-2` повинен мати властивість `flex-shrink` зі значенням `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
