---
id: 587d78ab367417b2b2512af2
title: "Використовуйте властивість flex-direction, щоб створити рядок"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

Додавання `display: flex` до елемента перетворює його на контейнер flex. Це дає змогу вирівнювати будь-які дочірні елементи цього елемента у рядки або стовпці. Ви робите це, додавши властивість `flex-direction` до батьківського елемента і встановивши її в row або column. Створення рядка вирівнює дочірні елементи по горизонталі, а створення стовпця — по вертикалі.

Інші варіанти для `flex-direction` — це `row-reverse` і `column-reverse`.

**Примітка:** Значення за замовчуванням для властивості `flex-direction` — `row`.

# --instructions--

Додайте CSS-властивість `flex-direction` до елемента `#box-container` і встановіть її значення в `row-reverse`.

# --hints--

Елемент `#box-container` має мати властивість `flex-direction`, встановлену в `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
