---
id: 587d78ae367417b2b2512afe
title: Використання скороченої властивості flex
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

Існує скорочений спосіб встановити кілька властивостей flex одночасно. Властивості `flex-grow`, `flex-shrink` і `flex-basis` можна встановити разом, використовуючи властивість `flex`.

Наприклад, `flex: 1 0 10px;` встановить елементу значення `flex-grow: 1;`, `flex-shrink: 0;` і `flex-basis: 10px;`.

Властивості за замовчуванням мають значення `flex: 0 1 auto;`.

# --instructions--

Додайте CSS-властивість `flex` до обох `#box-1` і `#box-2`. Задайте `#box-1` такі значення, щоб її `flex-grow` було `2`, її `flex-shrink` було `2`, а її `flex-basis` було `150px`. Задайте `#box-2` такі значення, щоб її `flex-grow` було `1`, її `flex-shrink` було `1`, а її `flex-basis` було `150px`.

Ці значення спричинять, що `#box-1` буде розширюватися, заповнюючи додатковий простір у два рази швидше, ніж `#box-2`, коли контейнер більший за 300px, і стискатися у два рази швидше, ніж `#box-2`, коли контейнер менший за 300px. 300px — це сумарний розмір значень `flex-basis` двох блоків.

# --hints--

Елемент `#box-1` повинен мати властивість `flex` зі значенням `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

Елемент `#box-2` повинен мати властивість `flex` зі значенням `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Ваш код має використовувати властивість `flex` для `#box-1` і `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
