---
id: 587d78ae367417b2b2512afc
title: Використання властивості flex-grow для розширення елементів
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

Протилежністю `flex-shrink` є властивість `flex-grow`. Нагадаємо, що `flex-shrink` контролює розмір елементів, коли контейнер стискається. Властивість `flex-grow` контролює розмір елементів, коли батьківський контейнер розширюється.

Використовуючи подібний приклад із попереднього завдання, якщо один елемент має значення `flex-grow` рівне `1`, а інший має значення `flex-grow` рівне `3`, то той, у кого значення `3`, зросте втричі більше за інший.

# --instructions--

Додайте CSS-властивість `flex-grow` до обох `#box-1` і `#box-2`. Дайте `#box-1` значення `1`, а `#box-2` — значення `2`.

# --hints--

Елемент `#box-1` повинен мати властивість `flex-grow` зі значенням `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

Елемент `#box-2` повинен мати властивість `flex-grow` зі значенням `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
