---
id: 587d78ae367417b2b2512aff
title: Використання властивості order для перестановки елементів
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

Властивість `order` використовується, щоб повідомити CSS порядок, у якому гнучкі елементи з’являються в гнучкому контейнері. За замовчуванням елементи відображаються в тому ж порядку, в якому вони йдуть у вихідному HTML. Ця властивість приймає числа як значення, і можна використовувати від’ємні числа.

# --instructions--

Додайте CSS-властивість `order` до обох `#box-1` і `#box-2`. Присвойте `#box-1` значення `2`, а `#box-2` — значення `1`.

# --hints--

Елемент `#box-1` має мати властивість `order` зі значенням `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

Елемент `#box-2` має мати властивість `order` зі значенням `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
