---
id: 587d78ad367417b2b2512afa
title: Використання властивості flex-wrap для перенесення рядка або колонки
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS гнучкий контейнер має функцію розбивати гнучкий контейнер на кілька рядків (або колонок). За замовчуванням гнучкий контейнер розміщує всі гнучкі елементи разом. Наприклад, рядок буде на одній лінії.

Однак використання властивості `flex-wrap` повідомляє CSS, що потрібно переносити елементи. Це означає, що зайві елементи переходять у новий рядок або колонку. Точка розриву, де відбувається перенесення, залежить від розміру елементів і розміру контейнера.

CSS також має варіанти напрямку перенесення:

<ul><li><code>nowrap</code>: це налаштування за замовчуванням, і елементи не переносяться.</li><li><code>wrap</code>: переносить елементи на кілька ліній зверху вниз, якщо це рядки, і зліва направо, якщо це колонки.</li><li><code>wrap-reverse</code>: переносить елементи на кілька ліній знизу вгору, якщо це рядки, і справа наліво, якщо це колонки.</li></ul>

# --instructions--

Поточний макет має забагато блоків для одного рядка. Додайте CSS-властивість `flex-wrap` до елемента `#box-container` і встановіть їй значення `wrap`.

# --hints--

Елемент `#box-container` повинен мати властивість `flex-wrap` зі значенням `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
