---
id: 587d78ab367417b2b2512af0
title: 'Використання display: flex для позиціювання двох блоків'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

У цьому розділі використовується чергування стилів завдань, щоб показати, як за допомогою CSS позиціювати елементи гнучким способом. Спочатку завдання пояснить теорію, а потім практичне завдання з простим компонентом твіта застосує концепцію гнучкого контейнера.

Додавання CSS-властивості `display: flex;` до елемента дозволяє використовувати інші властивості flex для створення адаптивної сторінки.

# --instructions--

Додайте CSS-властивість `display` до `#box-container` і встановіть її значення на `flex`.

# --hints--

`#box-container` має мати властивість `display` зі значенням `flex`.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
