---
id: 587d78af367417b2b2512b00
title: Використання властивості align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

Фінальна властивість для гнучких елементів — це `align-self`. Ця властивість дозволяє вам індивідуально налаштовувати вирівнювання кожного елемента, замість того, щоб встановлювати його для всіх одразу. Це корисно, оскільки інші загальні техніки правки за допомогою CSS-властивостей `float`, `clear` і `vertical-align` не працюють на гнучких елементах.

`align-self` приймає ті ж значення, що й `align-items`, і заміщатиме будь-яке значення, встановлене властивістю `align-items`.

# --instructions--

Додайте CSS-властивість `align-self` до обох `#box-1` і `#box-2`. Присвойте `#box-1` значення `center`, а `#box-2` — значення `flex-end`.

# --hints--

Елемент `#box-1` має мати властивість `align-self` зі значенням `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

Елемент `#box-2` має мати властивість `align-self` зі значенням `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
