---
id: a5229172f011153519423690
title: Сума всіх непарних чисел Фібоначчі
challengeType: 1
forumTopicId: 16084
dashedName: sum-all-odd-fibonacci-numbers
---

# --description--

Дано додатне ціле число `num`. Поверніть суму всіх непарних чисел Фібоначчі, які менші або дорівнюють `num`.

Перші два числа в послідовності Фібоначчі — 0 і 1. Кожне наступне число в послідовності є сумою двох попередніх чисел. Перші сім чисел послідовності Фібоначчі: 0, 1, 1, 2, 3, 5 та 8.

Наприклад, `sumFibs(10)` має повернути `10`, тому що всі непарні числа Фібоначчі, менші або рівні `10`, це 1, 1, 3 і 5.

# --hints--

`sumFibs(1)` має повернути число.

```js
assert(typeof sumFibs(1) === 'number');
```

`sumFibs(1000)` має повернути 1785.

```js
assert(sumFibs(1000) === 1785);
```

`sumFibs(4000000)` має повернути 4613732.

```js
assert(sumFibs(4000000) === 4613732);
```

`sumFibs(4)` має повернути 5.

```js
assert(sumFibs(4) === 5);
```

`sumFibs(75024)` має повернути 60696.

```js
assert(sumFibs(75024) === 60696);
```

`sumFibs(75025)` має повернути 135721.

```js
assert(sumFibs(75025) === 135721);
```

# --seed--

## --seed-contents--

```js
function sumFibs(num) {
  return num;
}

sumFibs(4);
```

# --solutions--

```js
function sumFibs(num) {
  var a = 1;
  var b = 1;
  var s = 0;
  while (a <= num) {
    if (a % 2 !== 0) {
      s += a;
    }
    a = [b, b=b+a][0];
  }
  return s;
}
```
