---
id: aa7697ea2477d1316795783b
title: свинська латина
challengeType: 1
forumTopicId: 16039
dashedName: pig-latin
---

# --description--

Свинська латина — це спосіб змінити англійські слова. Правила такі:

\- Якщо слово починається з приголосної, візьміть першу приголосну або групу приголосних, перенесіть її в кінець слова і додайте `ay`.

\- Якщо слово починається з голосної, просто додайте `way` у кінець.

# --instructions--

Перекладіть наданий рядок у свинську латину. Вхідні рядки гарантовано є англійськими словами в нижньому регістрі.

# --hints--

`translatePigLatin("california")` має повернути рядок `aliforniacay`.

```js
assert.deepEqual(translatePigLatin('california'), 'aliforniacay');
```

`translatePigLatin("paragraphs")` має повернути рядок `aragraphspay`.

```js
assert.deepEqual(translatePigLatin('paragraphs'), 'aragraphspay');
```

`translatePigLatin("glove")` має повернути рядок `oveglay`.

```js
assert.deepEqual(translatePigLatin('glove'), 'oveglay');
```

`translatePigLatin("algorithm")` має повернути рядок `algorithmway`.

```js
assert.deepEqual(translatePigLatin('algorithm'), 'algorithmway');
```

`translatePigLatin("eight")` має повернути рядок `eightway`.

```js
assert.deepEqual(translatePigLatin('eight'), 'eightway');
```

Потрібно обробляти слова, де перша голосна з’являється посередині слова. `translatePigLatin("schwartz")` має повернути рядок `artzschway`.

```js
assert.deepEqual(translatePigLatin('schwartz'), 'artzschway');
```

Потрібно обробляти слова без голосних. `translatePigLatin("rhythm")` має повернути рядок `rhythmay`.

```js
assert.deepEqual(translatePigLatin('rhythm'), 'rhythmay');
```

# --seed--

## --seed-contents--

```js
function translatePigLatin(str) {
  return str;
}

translatePigLatin("consonant");
```

# --solutions--

```js
function translatePigLatin(str) {
  if (isVowel(str.charAt(0))) return str + "way";
  var front = [];
  str = str.split('');
  while (str.length && !isVowel(str[0])) {
    front.push(str.shift());
  }
  return [].concat(str, front).join('') + 'ay';
}

function isVowel(c) {
  return ['a', 'e', 'i', 'o', 'u'].indexOf(c.toLowerCase()) !== -1;
}
```
