---
id: 68e6aff5325410118e76e7be
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过创建一个可访问的音频控件来练习使用 ARIA。

首先，为你的页面添加一个带有文本 `Audio Controls` 的 `h1` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应位于你的 `body` 元素内。

```js
assert.exists(document.querySelector('body > h1'));
```

你的 `h1` 元素应包含文本 `Audio Controls`。

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
