---
id: 68eb7f2397da251a40e04aa2
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

接下来，添加一个 `div` 元素来包含卷控件的所有元素。

在 `div` 元素内，嵌套一个 `span` 元素，其 `id` 属性设置为 `volume-label`。同时，`span` 元素的文本应为 `Volume`。

`id` 将在稍后被 ARIA 属性引用，以便辅助技术能够识别此文本作为卷控件的标签。

# --hints--

你应该在你的 `button` 元素下方有一个 `div` 元素。

```js
assert.exists(document.querySelector('button + div'));
```

你的 `div` 元素内应包含一个 `span` 元素。

```js
assert.exists(document.querySelector('button + div > span'));
```

你的 `span` 元素应具有 `id` 属性。

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

你的 `span` 元素应具有一个 `id` 属性，且其值设置为 `volume-label`。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

你的 `span` 元素应包含文本 `Volume`。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
