---
id: 68ec0a48c9d6db3594d768c3
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

现在是时候为你的页面添加另一个 `span` 元素。

此 `span` 元素应具有设置为 `volume-description` 的 `id` 属性。 此外，此 `span` 元素的文本应为 `Adjust the sound level`。

类似于之前的 `span` 元素，这个 `span` 的 `id` 将在稍后被 ARIA 属性引用，以便辅助技术能够识别此文本为卷控件的描述。

# --hints--

你应该在 `div` 元素内有两个 `span` 元素。

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

你的第二个 `span` 元素应具有 `id` 属性。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

你的第二个 `span` 元素应具有值为 `volume-description` 的 `id` 属性。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

你的第二个 `span` 元素应包含文本 `Adjust the sound level`。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
