---
id: 68ed488821845990c911a223
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

你的可访问音频控件几乎完成，但还有最后一件事需要添加。

在 `div` 元素下方添加一个 `button` 元素，其 `type` 属性设置为 `button`，并且文本为 `Mute`。

有了这个更改，你已经完成了本次研讨会！

# --hints--

你应该在 `div` 元素下方有一个 `button` 元素。

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

你的 `button` 元素应具有 `type` 属性。

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

你的 `button` 元素应具有值为 `button` 的 `type` 属性。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

你的 `button` 元素应有 `Mute` 文本。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
