---
id: 68ed4204ab7edc66b6cb2383
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

正如你从之前的课程中回忆的，`aria-labelledby` 属性用于页面上已有可用作标签的文本。在这种情况下，你已为你的 `span` 元素添加了 `id` 属性，该属性现在将在这里使用。

为你的 `input` 元素添加一个设置为 `volume-label volume-description` 的 `aria-labelledby` 属性。

# --hints--

你的 `input` 元素应具有 `aria-labelledby` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

你的 `input` 元素应具有值为 `volume-label volume-description` 的 `aria-labelledby` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
