---
id: 68ed3a96e0ebcc36565b39e5
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

接下来，你需要添加一个 `input` 元素，并将 `type` 属性设置为 `range` 来创建卷滑块。将 `min` 属性设置为 `0`，将 `max` 属性设置为 `100`，并将 `value` 属性设置为 `50` 以定义默认卷级别。

# --hints--

你应该在第二个 `span` 元素下方有一个 `input` 元素。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

你的 `input` 元素应具有 `type` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

你的 `input` 元素应具有值为 `range` 的 `type` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

你的 `input` 元素应具有 `min` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

你的 `input` 元素应具有值为 `0` 的 `min` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

你的 `input` 元素应具有 `max` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

你的 `input` 元素应具有值为 `100` 的 `max` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

你的 `input` 元素应具有 `value` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

你的 `input` 元素应具有值为 `50` 的 `value` 属性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
