---
id: 65f569725359e10d345bc52a
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Você vai construir um programa que calcula e desenha a trajetória de um projétil dado o ângulo, a velocidade e a altura do lançamento.

Comece importando `math`, você vai usá-lo muito neste projeto pois ele tem métodos úteis como `math.radians`, `math.cos`, `math.sin` e outros.

Também crie essas variáveis para terem o valor da aceleração gravitacional e alguns símbolos especiais que serão úteis mais tarde (use copiar e colar para estes).

```py
GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
```

# --hints--

Use `import math` para importar o módulo `math`.

```js
({test: () => assert(runPython(`_Node(_code).has_import('import math')`))})
```

Você deve ter uma variável `GRAVITATIONAL_ACCELERATION` com valor de `9.81`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("GRAVITATIONAL_ACCELERATION"), "missing variable"
assert _Node(_code).find_variable("GRAVITATIONAL_ACCELERATION").is_equivalent('GRAVITATIONAL_ACCELERATION = 9.81'), "variable has wrong value"
`)})
```

Você deve ter uma variável `PROJECTILE` com valor `"∙"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("PROJECTILE"), "missing variable"
assert _Node(_code).find_variable("PROJECTILE").is_equivalent('PROJECTILE = "∙"'), "variable has wrong value"
`)})
```

Você deve ter uma variável `x_axis_tick` com valor `"T"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("x_axis_tick"), "missing variable"
assert _Node(_code).find_variable("x_axis_tick").is_equivalent('x_axis_tick = "T"'), "variable has wrong value"
`)})
```

Você deve ter uma variável `y_axis_tick` com valor `"⊣"`.

```js
({test: () => runPython(`
assert _Node(_code).has_variable("y_axis_tick"), "missing variable"
assert _Node(_code).find_variable("y_axis_tick").is_equivalent('y_axis_tick = "⊣"'), "variable has wrong value"
`)})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
