---
id: 65fd4bd84561a14d3e720061
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

A variável de classe `__slots__` tem um uso especial em classes Python.Declarar `__slots__` e atribuí-lo a uma sequência de strings restringe a criação de atributos àqueles incluídos nessa sequência.Além disso, ele previne a criação do atributo especial `__dict__` e permite um acesso mais eficiente aos atributos.

Você deve usar a variável `__slots__` dentro da classe para definir quais atributos a classe possui: atribua a `__slots__` uma tupla contendo 3 strings, cada uma igual a um dos nomes dos atributos definidos no `__init__`.

# --hints--

Você deve atribuir a `__slots__` uma tupla que contenha `'__height'`, `'__speed'` e `'__angle'`.

```js
({test: () => assert(runPython(`
from itertools import permutations
slots = _Node(_code).find_class("Projectile").find_body().find_variable("__slots__")
perms = permutations(("__height", "__speed", "__angle"))
any(slots.is_equivalent(f'__slots__ = {perm}') for perm in perms)
`))})

```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--
class Projectile:

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
--fcc-editable-region--
```
