---
id: 65fd4b3bedee044b4b957d33
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Crie uma classe `Projectile` com um método `__init__` que inicializa a classe usando três argumentos: a velocidade inicial, a altura inicial e o ângulo inicial do lançamento do projétil, nesta ordem. 

Dentro do método `__init__`, atribua esses argumentos a três atributos privados: `__speed`, `__height` e `__angle`.O ângulo fornecido estará em graus; no entanto, ele deve ser armazenado internamente em radianos.Para conseguir isso, use a função `math.radians()` para converter o ângulo de graus para radianos ao atribuí-lo a `__angle`. 

O uso de dois underscores antes do nome de um atributo ativa o name mangling em Python.Isso significa que os atributos não são diretamente acessíveis de fora da classe usando seus nomes dados e devem ser acessados com os nomes modificados como `ball._Projectile__height` se necessário externamente, indicando que estes são destinados apenas para uso interno.

# --hints--

Você deve usar `class Projectile:` para criar a classe.

```js
({test: () => assert(runPython(`_Node(_code).has_class("Projectile")`))})

```

Você deve ter um método `__init__` dentro da classe `Projectile`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_body().has_function("__init__")`))})
```

Uma nova instância da classe `Projectile` deve ter os valores dos três atributos devidamente inicializados.

```js
({test: () => assert(runPython(`
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"
p = Projectile(20, 21, 22)
p._Projectile__height == 21 and p._Projectile__speed == 20 and p._Projectile__angle == math.radians(22)
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

--fcc-editable-region--

--fcc-editable-region--
```
