---
id: 660400f02cf3a76795b82d4e
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

A primeira coisa a configurar é um método que calcule o deslocamento do projétil, que é o espaço horizontal percorrido desde o lançamento até o momento em que o projétil toca o solo.

Crie um método `__calculate_displacement`, que tenha apenas `self` como parâmetro, e retorne o deslocamento do projétil.

Use a seguinte fórmula para calcular o deslocamento do projétil:  
\[ d = \frac{v \cdot \cos\theta \cdot \left(v\cdot\sin\theta + \sqrt{v^2 \cdot \sin^2\theta + 2 \cdot g \cdot h}\right)}{g} \]

Em que $d$ é o deslocamento, $v$ é a velocidade inicial, $\theta$ é o ângulo e $h$ é a altura inicial do projétil.Para $g$ você pode usar a variável `GRAVITATIONAL_ACCELERATION`.

Você deve usar os métodos `math.cos()` e `math.sin()` para as funções trigonométricas e `math.sqrt()` para calcular a raiz quadrada.Além disso, você deve saber que $x^y$ é escrito como `x ** y` em python.Também $\sin^2\theta$ significa que o valor resultante do seno é então elevado ao quadrado.

Lembre-se que com name mangling você precisa chamar o método como `_Projectile__calculate_displacement` se quiser testar, ou usá-lo de fora da classe:

```py
ball = Projectile(10, 3, 45)
displacement_of_ball = ball._Projectile__calculate_displacement() # 12.6173996009878
```

# --hints--

Você deve declarar um método chamado `__calculate_displacement` com `def __calculate_displacement(self):`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").has_function("__calculate_displacement")`))})
```

O método `__calculate_displacement` deve ter apenas o argumento `self`.

```js
({test: () => assert(runPython(`_Node(_code).find_class("Projectile").find_function("__calculate_displacement").has_args('self')`))})
```

O método `__calculate_displacement` deve retornar o valor correto.

```js
({test: () => assert(runPython(`
p = Projectile(20, 21, 22)
disp = p._Projectile__calculate_displacement()
round(disp, 2) == 55.06 and round(disp, 2) != disp
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
--fcc-editable-region--
    
--fcc-editable-region--
```
