---
id: 660fcf3e1b9bb056b2edb567
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Neste ponto você está pronto para criar a representação em string. Comece criando uma instância da classe `Projectile`. Abaixo da definição da classe, crie uma variável `ball` e atribua a ela uma chamada para `Projectile` usando `10, 3, 45` como argumentos.

Defina um método `__str__` dentro da classe, para dar à sua instância uma representação adequada em string, e faça-o retornar uma string com o seguinte formato:

```py

Projectile details:
speed: 10 m/s
height: 3 m
angle: 45°
displacement: 12.6 m

```

Deve começar e terminar com um caractere de nova linha, o ângulo deve ser escrito como um número inteiro em graus e o deslocamento deve ser impresso com uma casa decimal.

Você encontrará útil `math.degrees` para converter o ângulo de radianos para graus.

Quando estiver pronto você pode `print(ball)` para testar sua função.

# --hints--

Você deve ter uma variável `ball`.

```js
({test: () => assert(runPython(`_Node(_code).has_variable('ball')`))})
```

A variável `ball` deve ter um valor de `Projectile(10, 3, 45)`.

```js
({test: () => assert(runPython(`_Node(_code).find_variable('ball').is_equivalent('ball = Projectile(10, 3, 45)')`))})
```

A representação em string para `Projectile(45, 45, 45)` deve estar correta.

```js
({test: () => assert(runPython(
`
ball = Projectile(45, 45, 45)
str(ball) == """
Projectile details:
speed: 45 m/s
height: 45 m
angle: 45°
displacement: 244.4 m
"""`
))})
```

A representação em string também deve estar correta para outras instâncias.

```js
({test: () => assert(runPython(`p = Projectile(10, 10, 10)
str(
    p
) == """
Projectile details:
speed: 10 m/s
height: 10 m
angle: 10°
displacement: 15.9 m
"""`))})
```

Você deve ter uma chamada `print(ball)`.

```js
({
    test: () => runPython(`
    calls = _Node(_code).find_calls('print')
    assert any(c.is_equivalent('print(ball)') for c in calls)
    `)
})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)


    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
--fcc-editable-region--

--fcc-editable-region--
        
```
