---
id: 6616ec0bf9ee8548ce6b0f08
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Agora é hora de trabalhar no cálculo das coordenadas da trajetória!Crie um método da classe `Projectile` chamado `__calculate_y_coordinate`, ele deve ter, além de `self`, um parâmetro `x`.

\[ y = y_0 + x \tan\theta - \frac{g x^2}{2 v_0^2 \cos^2\theta} \]

A fórmula acima é para calcular a posição vertical $y$ para qualquer posição horizontal $x$, tendo o ângulo inicial $\theta$, velocidade $v_0$ e altura $y_0$.

Você precisará usar `math.tan()` e `math.cos()` e lembrar que `x ** y` é a forma de escrever $x^y$ e que o valor de $g$ está na variável `GRAVITATIONAL_ACCELERATION`.

Implemente o método para que ele retorne a coordenada $y$.

Se você quiser testar o método, de fora da classe, pode chamar `ball._Projectile__calculate_y_coordinate()` com um número como argumento depois que `ball` for declarado.

# --hints--

Sua classe `Projectile` deve ter um método chamado `__calculate_y_coordinate`.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').has_function('__calculate_y_coordinate')`))}})
```

O método `__calculate_y_coordinate` deve ter `self, x` como argumentos.

```js
({test: () => {assert(runPython(`_Node(_code).find_class('Projectile').find_function('__calculate_y_coordinate').has_args('self, x')`))}})
```

O método deve retornar o valor correto. Para `Projectile(12, 12, 12)`, dadas as 4 coordenadas $x$ de 0, 1, 2 e 3, a coordenada y deve ser aproximadamente 12, 12.18, 12.28 e 12.32.

```js
({test: () => {
    assert(runPython(`
    a = Projectile(12, 12, 12)
    all([round(a._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(0, 12.0), (1, 12.18), (2, 12.28), (3, 12.32)]])
    `))
}})
```

O método deve retornar o valor correto. Para `Projectile(45, 12, 22)`, dadas as coordenadas $x$ de 6, 9, 10 e 12 a coordenada y deve ser aproximadamente 14.32, 15.41, 15.76, 16.44.

```js
({test: () => {
    assert(runPython(`
    b = Projectile(45, 12, 22)
    all([round(b._Projectile__calculate_y_coordinate(x), 2) == y for x,y in [(6, 14.32), (9, 15.41), (10, 15.76), (12, 16.44)]])
    `))
}})
```

O método `__calculate_y_coordinate` não deve arredondar o resultado.

```js
({test: () => assert.isFalse(runPython(`Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5) == round(Projectile(12, 13, 14)._Projectile__calculate_y_coordinate(5), 2)`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    --fcc-editable-region--
    
    --fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
   
```
