---
id: 662fa453ec8033322b3ab2e3
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

É hora de falar sobre encapsulamento e getters. Você escreveu os três atributos de instância para serem privados usando um duplo sublinhado inicial. Observe que esses atributos são chamados de privados por convenção: embora ainda possam ser acessados de fora, é acordado que não se deve fazer isso.

Getters são o que pode ser usado para obter os valores de fora. Para definir um getter, você define um método que retorna o valor do atributo desejado e o decora com `@property`:

```py
class Nest:
    ...
    @property
    def number_of_eggs(self):
        return self.__number_of_eggs
```

 O decorator transforma o método em uma propriedade, o que significa que o método não é chamado como um método comum, mas é usado como um atributo:

```py
n = Nest()
print(n.number_of_eggs)
```

No exemplo acima, o atributo privado `__number of eggs` é acessado através da propriedade `number_of_eggs` de `n`.

Crie um getter chamado `speed` e faça com que ele retorne o valor do atributo privado `__speed`.

# --hints--

Você deve definir um método `speed`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('speed')
`))})
```

Seu método `speed` deve ter um único parâmetro, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_args('self')
`))})
```

Seu método `speed` deve ter um decorador `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('speed').has_decorators('property')
`))})
```

Seu método `speed` deve retornar o valor do atributo privado `__speed`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.speed == 22
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()
   
```
