---
id: 6634fa2417d7194b0c9425f9
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Agora crie os dois getters para `__height` e `__angle`.Lembre-se que para o ângulo, o valor é dado para a classe em graus, mas é salvo internamente em radianos, então faça o getter retornar o valor em graus.Você pode usar `math.degrees` e a função `round` para isso.

# --hints--
Você deve definir um método `height`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('height')
`))})
```

Seu método `height` deve ter um único parâmetro, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_args('self')
`))})
```

Seu método `height` deve ter um decorador `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('height').has_decorators('property')
`))})
```

Seu método `height` deve retornar o valor do atributo privado `__height`.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.height == 23
`))})
```

Você deve definir um método `angle`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').has_function('angle')
`))})
```

Seu método `angle` deve ter um decorador `@property`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_decorators('property')
`))})
```

Seu método `angle` deve ter um único parâmetro, `self`.

```js
({test: () => assert(runPython(`
_Node(_code).find_class('Projectile').find_function('angle').has_args('self')
`))})
```

Seu método `angle` deve retornar o valor do atributo privado `__angle` convertido para graus.

```js
({test: () => assert(runPython(`
a = Projectile(22, 23, 24)
a.angle == 24
`))})
```

# --seed--

## --seed-contents--

```py
import math

GRAVITATIONAL_ACCELERATION = 9.81
PROJECTILE = "∙"
x_axis_tick = "T"
y_axis_tick = "⊣"

class Projectile:
    __slots__ = ('__speed', '__height', '__angle')

    def __init__(self, speed, height, angle):
        self.__speed = speed
        self.__height = height
        self.__angle = math.radians(angle)
        
    def __str__(self):
        return f'''
Projectile details:
speed: {self.__speed} m/s
height: {self.__height} m
angle: {round(math.degrees(self.__angle))}°
displacement: {round(self.__calculate_displacement(), 1)} m
'''

    def __calculate_displacement(self):
        horizontal_component = self.__speed * math.cos(self.__angle)
        vertical_component = self.__speed * math.sin(self.__angle)
        squared_component = vertical_component**2
        gh_component = 2 * GRAVITATIONAL_ACCELERATION * self.__height
        sqrt_component = math.sqrt(squared_component + gh_component)
        
        return horizontal_component * (vertical_component + sqrt_component) / GRAVITATIONAL_ACCELERATION
        
    def __calculate_y_coordinate(self, x):
        height_component = self.__height
        angle_component = math.tan(self.__angle) * x
        acceleration_component = GRAVITATIONAL_ACCELERATION * x ** 2 / (
                2 * self.__speed ** 2 * math.cos(self.__angle) ** 2)
        y_coordinate = height_component + angle_component - acceleration_component

        return y_coordinate
    
    def calculate_all_coordinates(self):
        return [
            (x, self.__calculate_y_coordinate(x))
            for x in range(math.ceil(self.__calculate_displacement()))
        ]

--fcc-editable-region--
    @property
    def speed(self):
        return self.__speed

    
--fcc-editable-region--

ball = Projectile(10, 3, 45)
print(ball)
coordinates = ball.calculate_all_coordinates()

   
```
