---
id: 66b62d0ad68488dd76228d6c
title: Відлагодження випадкового змінювача фону кольору
challengeType: 25
dashedName: debug-a-random-background-color-changer
demoType: onClick
---


# --description--

Camperbot вивчає JavaScript і намагався створити власний випадковий перемикач кольору фону. Однак він припустився кількох помилок.

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб лабораторна робота працювала правильно.

**Історії користувача:**

1. Ви повинні виправити синтаксичні помилки в масиві `darkColorsArr`.
1. Ви повинні виправити `ReferenceError`, що виникає при виклику об’єкта `math` всередині функції `getRandomIndex`.
1. Ви повинні округлити `randomIndex` до найближчого цілого числа вниз.
1. Ви повинні оновити значення змінної `body`, щоб використовувати правильну назву методу, і виправити пов’язану `TypeError`.
1. Ви повинні виправити `bgHexCodeSpanElement`, щоб він правильно вибирав елемент із ідентифікатором `id` `bg-hex-code`.
1. Ви повинні змінити змінну `color` всередині `changeBackgroundColor`, щоб вона встановлювалася у випадковий колір із масиву `darkColorsArr`.
1. Ви повинні оновити змінну `btn`, щоб вона цілеспрямовано вибирала кнопку. Відкрийте `index.html`, щоб побачити правильну назву `id` для цієї кнопки.
1. Ви повинні виправити рядок, який додає слухача події для події `click` до `btn`, передаючи референс функції `changeBackgroundColor` методу.
1. Ваш фон повинен випадково змінюватися на один із кольорів у `darkColorsArr` щоразу, коли натискають на `#btn`.

# --hints--

Ви повинні виправити синтаксичні помилки в масиві `darkColorsArr`. Пам’ятайте, що елементи масиву мають розділятися комами.

```js
assert.isArray(darkColorsArr);
```

Ви повинні виправити помилку регістру в рядку `math.random()`.

```js
assert.match(getRandomIndex.toString(), /\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*/);
```

Ви повинні округлити `darkColorsArr.length * Math.random()` вниз до найближчого цілого числа.

```js
assert.match(getRandomIndex.toString(), /Math\.floor\(\s*(darkColorsArr\.length\s*\*\s*Math\.random\(\s*\)|Math\.random\(\s*\)\s*\*\s*darkColorsArr\.length)\s*\)/);
```

Ви повинні оновити значення змінної `body`, щоб `queryselector` використовував camel case.

```js
assert.match(code, /const\s+body\s*=\s*document\.querySelector\s*\(\s*("|')body\1\s*\)/);
```

Ви повинні виправити `"bg-hex-code"`, щоб він правильно посилався на ідентифікатор `id` об’єкта `bg-hex-code`. Пам’ятайте, що селектори `id` у CSS починаються з символу решітки (`#`).

```js
assert.match(bgHexCodeSpanElement.id, /bg-hex-code/);
```

Ви повинні викликати функцію `getRandomIndex` у рядку `darkColorsArr[getRandomIndex]`.

```js
assert.match(changeBackgroundColor.toString(), /darkColorsArr\[\s*getRandomIndex\(\s*\)\s*\]/);
```

Ви повинні виправити назву `id` у рядку `"#click-btn"`, щоб вона відповідала правильній назві методу `id` у файлі `index.html`.

```js
assert.match(btn?.id, /btn/);
```

Ви не повинні викликати функцію `changeBackgroundColor`. Замість цього передайте референс функції методу `addEventListener`.

```js
assert.match(code, /btn\.addEventListener\s*\(\s*("|')click\1\s*,\s*changeBackgroundColor\s*\)/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
    "#2C3E50",
    "#34495E",
    "#2C2C2C"
    "#616A6B"
    "#4A235A"
    "#2F4F4F",
    "#0E4B5A",
    "#36454F",
    "#2C3E50",
    "#800020",
  ];
  
  function getRandomIndex() {
    const randomIndex = darkColorsArr.length * math.random();
    return randomIndex;
  }
  console.log(getRandomIndex())
  
  const body = document.queryselector("body");
  const bgHexCodeSpanElement = document.querySelector("bg-hex-code");
  console.log(bgHexCodeSpanElement)
  
  function changeBackgroundColor() {
    const color = darkColorsArr[getRandomIndex];
  
    bgHexCodeSpanElement.innerText = color;
    body.style.backgroundColor = color;
  }
  const btn = document.querySelector("#click-btn");
  console.log(btn)
  
  btn.addEventListener("click", changeBackgroundColor());
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Debug a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
</head>

<body>
    <h1>Random Background Color changer</h1>

    <main>
        <section class="bg-information-container">
            <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
        </section>

        <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
</body>

</html>
```

```css
*,
*::before,
*::after {
    margin: 0;
    padding: 0;
    box-sizing: border-box;
}

:root {
    --yellow: #f1be32;
    --golden-yellow: #feac32;
    --dark-purple: #110815;
    --light-grey: #efefef;
}

body {
    background-color: var(--dark-purple);
    color: var(--light-grey);
    text-align: center;
}

.bg-information-container {
    margin: 15px 0 25px;
    font-size: 1.2rem;
}

.btn {
    cursor: pointer;
    padding: 10px;
    margin: 10px;
    color: var(--dark-purple);
    background-color: var(--golden-yellow);
    background-image: linear-gradient(#fecc4c, #ffac33);
    border-color: var(--golden-yellow);
    border-width: 3px;
}

.btn:hover {
    background-image: linear-gradient(#ffcc4c, #f89808);
}
```

```js
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C",
  "#616A6B",
  "#4A235A",
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];

function getRandomIndex() {
  const randomIndex = Math.floor(darkColorsArr.length * Math.random());
  return randomIndex;
}

const body = document.querySelector("body");
const bgHexCodeSpanElement = document.querySelector("#bg-hex-code");

function changeBackgroundColor() {
  const color = darkColorsArr[getRandomIndex()];

  bgHexCodeSpanElement.innerText = color;
  body.style.backgroundColor = color;
}
const btn = document.querySelector("#btn");

btn.addEventListener("click", changeBackgroundColor);
```
