---
id: 6626a060c4006f793e10cb33
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

今度は、ある学生のスコアを A～F の成績に変換するのにあなたの助けが必要です。

数値 `score` を引数として受け取る `getGrade` 関数を完成させましょう。 関数は、スコアに基づいて成績を表す文字列を返す必要があります。

こちらがスコアとそれに対応する成績の一覧です:

| スコア範囲     | 成績      |
| --------- | ------- |
| `100`     | `"A++"` |
| `90 - 99` | `"A"`   |
| `80 - 89` | `"B"`   |
| `70 - 79` | `"C"`   |
| `60 - 69` | `"D"`   |
| `0 - 59`  | `"F"`   |


**ヒント**

- 条件文 (`if`, `else if`, `else`) について学んだことを思い出してください。
- 比較演算子 (`>`, `<`, `>=`, `<=`, `===`) について学んだことを思い出してください。
- Remember that you learned about the logical AND operator (`&&`).

# --hints--

`getGrade` 関数は、スコアが `100` の場合 `"A++"` を返すべきです。

```js
assert.strictEqual(getGrade(100), "A++");
```

`getGrade` 関数は、スコアが `94` の場合 `"A"` を返すべきです。

```js
assert.strictEqual(getGrade(94), "A");
```

`getGrade` 関数は、スコアが `80` から `89` の間の場合、`"B"` を返す必要があります。

```js
assert.strictEqual(getGrade(80), "B");
assert.strictEqual(getGrade(83), "B");
assert.strictEqual(getGrade(88), "B");
```

`getGrade` 関数は、スコアが `78` の場合、`"C"` を返す必要があります。

```js
assert.strictEqual(getGrade(75), "C");
```

`getGrade` 関数は、スコアが `60` から `69` の間の場合、`"D"` を返す必要があります。

```js
assert.strictEqual(getGrade(60), "D");
assert.strictEqual(getGrade(63), "D");
assert.strictEqual(getGrade(68), "D");
```

`getGrade` 関数は、スコアが `57` の場合、`"F"` を返す必要があります。

```js
assert.strictEqual(getGrade(57), "F");
```

`getGrade` 関数は、スコアが `0` から `59` の間の場合、`"F"` を返す必要があります。

```js
assert.strictEqual(getGrade(0), "F");
assert.strictEqual(getGrade(30), "F");
assert.strictEqual(getGrade(43), "F");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}
--fcc-editable-region--
function getGrade(score) {

}

console.log(getGrade(96));
console.log(getGrade(82));
console.log(getGrade(56));
--fcc-editable-region--
```
