---
id: 662693f82c91a66be46c881b
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

とある先生が学生たちのテストの採点を終え、クラスの平均点を計算するためにあなたの助けを必要としています。

`getAverage` 関数を完成させてください。この関数はスコア (テストの点数) の配列を受け取り、平均点を返します。

平均点は、点数の合計を、スコアの件数で割ることで計算できます。

```md
average = sum of all scores / total number of scores
```

コードの動作をテストできるように、いくつかの関数呼び出しが用意されています。

**ヒント**

- `scores` 配列を反復処理してすべての点数を合計するには、ループが使えます。
- スコアの件数を取得するには `length` プロパティが使えます。

# --hints--

`getAverage` 関数は数値を返す必要があります。

```js
assert.strictEqual(typeof getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 'number');
```

`getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89])` は `71.7` を返すべきです。

```js
assert.strictEqual(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]), 71.7);
```

`getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95])` は `85.4` を返すべきです。

```js
assert.strictEqual(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]), 85.4);
```

`getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100])` は `92.4` を返すべきです。

```js
assert.strictEqual(getAverage([38, 99, 87, 100, 100, 100, 100, 100, 100, 100]), 92.4);
```

`getAverage` 関数は任意のテストスコアの平均点を返す必要があります。

```js
assert.strictEqual(getAverage([52, 56, 60, 65, 70, 75, 80, 85, 90, 95]), 72.8);
assert.strictEqual(getAverage([45, 50, 55, 60, 65, 70, 75, 80, 85, 90]), 67.5);
assert.strictEqual(getAverage([38, 42, 46, 50, 54, 58, 62, 66, 70, 74]), 56);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function getAverage(scores) {

}

console.log(getAverage([92, 88, 12, 77, 57, 100, 67, 38, 97, 89]));
console.log(getAverage([45, 87, 98, 100, 86, 94, 67, 88, 94, 95]));
--fcc-editable-region--
```
