---
id: 6626b4c58c027d86478ff5eb
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

先生はあなたの作成したプログラムにとても満足しています。 そしてさらに、ある学生の成績が合格どうかを簡単に確認したいと考えています。 合格となるのは、`"F"` 以外です。

学生のスコアを引数として受け取る `hasPassingGrade` 関数を完成させてください。 合格の場合は `true` を、不合格の場合は `false` を返すように関数を作成しましょう。

**ヒント**

- `getGrade` 関数を使用して、スコアから成績を取得します。 次に、その成績が合格かどうかをチェックしてください。

# --hints--

`hasPassingGrade` 関数はブール値を返す必要があります。

```js
assert.strictEqual(typeof hasPassingGrade(100), "boolean");
```

`hasPassingGrade` 関数は、成績が `"A"` の場合には `true` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(100));
```

`hasPassingGrade` 関数は、成績が `"F"` の場合には `false` を返す必要があります。

```js
assert.isFalse(hasPassingGrade(53));
```

`hasPassingGrade` 関数は、合格となるすべての成績に対して `true` を返す必要があります。

```js
assert.isTrue(hasPassingGrade(87));
assert.isTrue(hasPassingGrade(60));
assert.isTrue(hasPassingGrade(73));
assert.isTrue(hasPassingGrade(96));
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

--fcc-editable-region--
function hasPassingGrade(score) {

}


console.log(hasPassingGrade(100));
console.log(hasPassingGrade(53));
console.log(hasPassingGrade(87));
--fcc-editable-region--
```
