---
id: 6626b8dcf5057f896f948440
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

先生はこれで必要なすべての情報を入手できたので、結果を学生に送信するメッセージを作成したいと考えています。

関数 `studentMsg` を完成させてください。この関数は、クラス全員のスコアの配列 `totalScores` と、ある学生のスコア `studentScore` を引数に取ります。 関数はその学生へのメッセージを表す文字列を返します。

その学生が合格の場合、次の形式の文字列を返してください。

```md
Class average: average-goes-here. Your grade: grade-goes-here. You passed the course.
```

その学生が不合格の場合、次の形式の文字列を返してください。

```md
Class average: average-goes-here. Your grade: grade-goes-here. You failed the course.
```

`average-goes-here` の部分を、クラスの平均点に置き換えてください。 `grade-goes-here` の部分を、その学生の成績に置き換えてください。

**ヒント**

- クラスの平均点を取得するには `getAverage` 関数を使ってください。
- 学生の成績を取得するには `getGrade` 関数を使ってください。
- メッセージを組み立てるには文字列連結 (`+`) を使ってください。
- メッセージの句読点やスペースの有無に注意してください。

# --hints--

`studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37)` の関数呼び出しは、次のメッセージを返すべきです: `"Class average: 71.7. Your grade: F. You failed the course."`.

```js
assert.strictEqual(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37), "Class average: 71.7. Your grade: F. You failed the course.");
```

`studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100)` の関数呼び出しは、次のメッセージを返すべきです: `"Class average: 50.8. Your grade: A++. You passed the course."`.

```js
assert.strictEqual(studentMsg([56, 23, 89, 42, 75, 11, 68, 34, 91, 19], 100), "Class average: 50.8. Your grade: A++. You passed the course.");
```

`studentMsg` 関数は、与えられた学生のスコアとクラスの平均点に基づいて正しいメッセージを返すべきです。

```js
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 92), "Class average: 70.25. Your grade: A. You passed the course.");
assert.strictEqual(studentMsg([33, 44, 55, 66, 77, 88, 99, 100], 57), "Class average: 70.25. Your grade: F. You failed the course.");
```

# --seed--

## --seed-contents--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

--fcc-editable-region--
function studentMsg(totalScores, studentScore) {

}
console.log(studentMsg([92, 88, 12, 77, 57, 100, 67, 38, 97, 89], 37));
--fcc-editable-region--

```


# --solutions--

```js
function getAverage(scores) {
  let sum = 0;

  for (const score of scores) {
    sum += score;
  }

  return sum / scores.length;
}

function getGrade(score) {
  if (score === 100) {
    return "A++";
  } else if (score >= 90) {
    return "A";
  } else if (score >= 80) {
    return "B";
  } else if (score >= 70) {
    return "C";
  } else if (score >= 60) {
    return "D";
  } else {
    return "F";
  }
}

function hasPassingGrade(score) {
  return getGrade(score) !== "F";
}

function studentMsg(totalScores, studentScore) {
  let average = getAverage(totalScores);
  let grade = getGrade(studentScore);

  return `Class average: ${average}. Your grade: ${grade}. You ${
    hasPassingGrade(studentScore) ? "passed" : "failed"
  } the course.`;
}

```
