---
id: 668f08ea07b99b1f4a91acab
title: Jenga ukurasa wa mapishi
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara.

**Hadithi za watumizi:**

1. Unapaswa kuwa na tamko la `!DOCTYPE html`.
1. Unapaswa kuwa na kipengele cha `html` chenye `lang` imewekwa kuwa `en`.
1. Unapaswa kuwa na kipengele cha `head` chenye kipengele cha `title` chenye jina la mapishi yako, na kipengele cha `meta` chenye sifa ya kipengele `charset` imewekwa kuwa `UTF-8`.
1. Unapaswa kuwa na kipengele cha `body`.
1. Unapaswa kuwa na kipengele kimoja tu cha `h1` chenye jina la mapishi yako.
1. Unapaswa kuwa na kipengele cha `p` kinachoanzisha mapishi chini ya `h1`.
1. Unapaswa kuwa na kipengele kimoja cha `h2` chenye maandishi `Ingredients` kwa sehemu ya viungo.
1. Unapaswa kuwa na orodha isiyopangwa (`ul` kipengele) yenye angalau vitu vinne vya orodha (`li` vipengele) vinavyoorodhesha viungo vyako chini ya kipengele cha kwanza cha `h2`.
1. Unapaswa kuwa na kipengele cha pili cha `h2` chenye maandishi `Instructions` kwa sehemu ya maelekezo.
1. Unapaswa kuwa na orodha iliyopangwa (`ol` kipengele) yenye angalau vitu vinne vya orodha vinavyoorodhesha hatua za mapishi kwa mpangilio, chini ya kipengele cha pili cha `h2`.
1. Unapaswa kuwa na kipengele kimoja cha `img` chenye sifa ya kipengele `src` imewekwa kwa picha halali na sifa ya kipengele `alt` inayoelezea picha. Unaweza kutumia `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg` ikiwa unataka.

# --hints--

Ukurasa wako wa mapishi unapaswa kuwa na tamko la `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE html>/i);
```

Unapaswa kuwa na kipengele cha `html` chenye `lang` kimesetwa kuwa `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

Unapaswa kuwa na kipengele cha `head` ndani ya kipengele cha `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Unapaswa kuwa na kipengele cha `title` ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Kipengele chako cha `title` kinapaswa kuwa na kichwa cha mapishi yako.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Unapaswa kuwa na kipengele cha `meta` ndani ya kipengele chako cha `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Kipengele chako cha `meta` kinapaswa kuwa na sifa ya kipengele ya `charset` iliyowekwa kuwa `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Unapaswa kuwa na kipengele cha `body` ndani ya kipengele chako cha `html`.

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

Unapaswa kuwa na kipengele kimoja tu cha `h1` chenye jina la mapishi yako.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Unapaswa kuwa na kipengele kimoja tu cha `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Unapaswa kuwa na kipengele cha `p` chini ya kipengele chako cha `h1`.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

Kipengele chako cha kwanza cha `p` kinapaswa kuelezea mapishi yako.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Kipengele chako cha kwanza cha `h2` kinapaswa kuwa na maandishi ya `Ingredients`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

Unapaswa kuwa na kipengele cha orodha isiyopangwa chini ya kipengele chako cha kwanza cha `h2`.

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

Unapaswa kuwa na angalau vitu vinne vya orodha katika orodha isiyopangwa yenye viungo.

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Kipengele chako cha pili cha `h2` kinapaswa kuwa na maandishi ya `Instructions`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

Unapaswa kuwa na kipengele cha orodha iliyopangwa chini ya kipengele chako cha pili cha `h2`.

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

Unapaswa kuwa na angalau vitu vinne vya orodha katika orodha iliyopangwa yenye maelekezo.

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

Unapaswa kuwa na angalau kipengele kimoja cha `img`.

```js
assert.exists(document.querySelector('img'));
```

Vipengele vyote vya `img` vinapaswa kuwa na sifa ya kipengele halali ya `src` na thamani.

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

Vipengele vyote vya `img` vinapaswa kuwa na sifa ya kipengele ya `alt` kuelezea picha.

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
