---
id: 68e42de6396f5c49d90d1951
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Em uma lição anterior, você aprendeu sobre o método `replace()` que é usado para localizar uma substring em uma string e substituí-la por outro valor.

Aqui está um exemplo:

```js
const text = "I love JavaScript!";
const newText = text.replace("JavaScript", "coding");
// Result: "I love coding!"
```

Lembre-se de que strings são imutáveis o que significa que o `text` original não é modificado nesse exemplo. Uma nova string é criada em vez disso. 

Crie uma nova variável chamada `replacedString` e atribua a ela `originalString.replace("cats", "dogs")`. 

# --hints--

Você deve ter uma variável chamada `replacedString`.

```js
assert.exists(replacedString);
```

Sua variável `replacedString` deve ser uma string.

```js
assert.isString(replacedString);
```

Você deve atribuir `originalString.replace("cats", "dogs")` à variável `replacedString`.

```js
assert.equal(replacedString, "I love dogs.");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

--fcc-editable-region--

--fcc-editable-region--
```
