---
id: 68e41084bb275539f72de9fe
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você praticará o uso dos métodos `replace()`, `replaceAll()` e `repeat()` construindo um aplicativo transformador de strings. 

Para começar, crie uma variável chamada `originalString` e atribua a ela o valor `"I love cats."`.

# --hints--

Você deve ter uma variável chamada `originalString`.

```js
assert.exists(originalString);
```

Sua variável `originalString` deve ser uma string.

```js
assert.isString(originalString);
```

Você deve atribuir a string `"I love cats."` à variável `originalString`.

```js
assert.equal(originalString, "I love cats.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
