---
id: 68e4417b7317835208145323
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Em seguida, crie uma nova variável chamada `dogsOnlySentence` e atribua a ela `exampleSentence.replace("cats", "dogs")`.

Então, abaixo daquela variável, registre no console `"Replacing all occurrences of cats with dogs:"`. Abaixo daquele `console.log()`, adicione outro `console.log()` para a variável `dogsOnlySentence`. 

Você deve agora ver o texto `"I love dogs and cats are so much fun!"` registrado no console. Mas esse não é o resultado desejado. Na próxima etapa, você aprenderá como corrigir isso para que o texto fique `"I love dogs and dogs are so much fun!"`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável `dogsOnlySentence`.

```js
assert.exists(dogsOnlySentence);
```

Sua variável `dogsOnlySentence` deve ser uma string.

```js
assert.isString(dogsOnlySentence);
```

Você deve atribuir `exampleSentence.replace("cats", "dogs")` à sua variável `dogsOnlySentence`.

```js
assert.equal(dogsOnlySentence, "I love dogs and cats are so much fun!");
```

Você deve registrar `"Replacing all occurrences of cats with dogs:"` no console.

```js
assert.equal(getLogs()[6], "Replacing all occurrences of cats with dogs:");
```

Você deve registrar a variável `dogsOnlySentence` no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingDogsOnlySentence = codeWithoutComments.match(/console\.log\(\s*dogsOnlySentence\s*\)/g)

assert.isAtLeast(loggingDogsOnlySentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--

--fcc-editable-region--
```
