---
id: 68e4465fba0bd256868b7deb
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Agora é hora de trabalhar com um novo método de string.

Comece criando uma variável chamada `learningSentence` e atribua a string `"I love learning!"`. 

Então, abaixo dessa variável, registre no console `"Original learning sentence:"`. Abaixo desse `console.log()`, adicione outro `console.log()` para a variável `learningSentence`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável `learningSentence`.

```js
assert.exists(learningSentence);
```

Sua variável `learningSentence` deve ser uma string.

```js
assert.isString(learningSentence);
```

Você deve atribuir a string `"I love learning!"` à variável `learningSentence`.

```js
assert.equal(learningSentence, "I love learning!");
```

Você deve registrar `"Original learning sentence:"` no console.

```js
assert.equal(getLogs()[8], "Original learning sentence:");
```

Você deve registrar a variável `learningSentence` no console.

```js
assert.equal(getLogs()[9], learningSentence);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

--fcc-editable-region--

--fcc-editable-region--
```
