---
id: 68e43f94365a964f26c88f78
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora é hora de trabalhar com um novo método de string. 

Comece criando uma nova variável chamada `exampleSentence` e atribua a ela a string `"I love cats and cats are so much fun!"`.

Então, abaixo daquela variável, registre no console `"Original sentence:"`. Então, abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `exampleSentence`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável `exampleSentence`.

```js
assert.exists(exampleSentence);
```

Sua variável `exampleSentence` deve ser uma string.

```js
assert.isString(exampleSentence);
```

Você deve atribuir a string `"I love cats and cats are so much fun!"` à variável `exampleSentence`.

```js
assert.equal(exampleSentence, "I love cats and cats are so much fun!");
```

Você deve registrar `"Original sentence:"` no console.

```js
assert.equal(getLogs()[4], "Original sentence:");
```

Você deve registrar a variável `exampleSentence` no console.

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingExSentence = codeWithoutComments.match(/console\.log\(\s*exampleSentence\s*\)/g)

assert.isAtLeast(loggingExSentence.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

--fcc-editable-region--

--fcc-editable-region--
```
