---
id: 68e44519c769bf545bb3f22a
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Para obter a saída desejada de `"I love dogs and dogs are so much fun!"`, você precisará usar o método `replaceAll()` em vez do método `replace()`. Este método retorna uma nova string para todas as correspondências da substring. 

Atualize o `exampleSentence.replace("cats", "dogs")` para usar o método `replaceAll()` em vez disso. Agora você deve ver o texto correto no console.

# --hints--

Você deve atualizar a linha de código `exampleSentence.replace("cats", "dogs")` para usar o método `replaceAll()`.

```js
assert.equal(dogsOnlySentence, "I love dogs and dogs are so much fun!");
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

--fcc-editable-region--
const dogsOnlySentence = exampleSentence.replace("cats", "dogs");
--fcc-editable-region--
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);
```
