---
id: 68e448ec4ac4205968c0a805
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Seria bom se a palavra `"love"` fosse repetida três vezes na frase `"I love learning!"`. 

Como você se lembra de uma lição anterior, você pode repetir uma string um número específico de vezes usando o método `repeat()`. Aqui está um exemplo:

```js
const word = "Hello!";
const repeatedWord = word.repeat(3);

console.log(repeatedWord);  // "Hello!Hello!Hello!"
```

Como as strings são imutáveis, este método não modificará a string original. Ele retornará uma nova string com o conteúdo repetido. 

Crie uma variável chamada `repeatedLove` e atribua `"love ".repeat(3)`. Depois, registre no console a variável `repeatedLove` para que você possa ver o resultado. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável chamada `repeatedLove`.

```js
assert.exists(repeatedLove);
```

Sua variável `repeatedLove` deve ser uma string.

```js
assert.isString(repeatedLove);
```

Você deve atribuir `"love ".repeat(3)` à variável `repeatedLove`.

```js
assert.equal(repeatedLove, "love ".repeat(3));
```

Você deve registrar a variável `repeatedLove` no console.

```js
assert.equal(getLogs()[10], repeatedLove);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

--fcc-editable-region--

--fcc-editable-region--
```
