---
id: 68e43e884ef09d4d2ab05631
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora é hora de ver o resultado de usar o método `replace()`.

Comece adicionando um `console.log()` com o valor `"After using the replace() method:"`.

Então abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `replacedString`.

Dê uma olhada no console e você deve ver que a nova frase diz `"I love dogs."` em vez da original `"I love cats."`.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar `"After using the replace() method:"` no console.

```js
assert.equal(getLogs()[2], "After using the replace() method:");
```

Você deve registrar a variável `replacedString` no console. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingReplacedStr = codeWithoutComments.match(/console\.log\(\s*replacedString\s*\)/g)

assert.isAtLeast(loggingReplacedStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
--fcc-editable-region--

--fcc-editable-region--
```
