---
id: 68e44abea678a65b877864dc
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Seria bom ver essa palavra repetida em uma frase. Crie uma nova variável chamada `newSentence` e atribua a ela o resultado da concatenação da palavra `"I "` com a variável `repeatedLove` seguida da concatenação da string `" learning."`. Em seguida, registre a variável `newSentence` no console. Você pode escolher usar o operador `+` ou template literals para a concatenação de strings. 

Agora você deve ver o texto `"I love love love  learning."`. 

*NOTA*: Há um espaço extra após o último `"love"` que está correto. Na última etapa, você irá remover esse espaço extra. 

# --hints--

Você deve ter uma variável chamada `newSentence`.

```js
assert.exists(newSentence);
```

Sua variável `newSentence` deve ser uma string.

```js
assert.isString(newSentence);
```

Você deve concatenar a palavra `"I "` com a variável `repeatedLove` seguida da concatenação da string `" learning."`. Esse resultado deve ser atribuído à variável `newSentence`. Tenha cuidado com os espaços. 

```js
assert.equal(newSentence, `I ${repeatedLove} learning.`);
```

Você NÃO deve atribuir a string literal `"I love love love  learning."` à variável `newSentence`. Por favor, releia as instruções para a abordagem correta.

```js
assert.notMatch(code, /(["'])I\s+love\s+love\s+love\s+\s+learning.(["'])/);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
console.log("Original string:");
console.log(originalString);

const replacedString = originalString.replace("cats", "dogs");
console.log("After using the replace() method:");
console.log(replacedString);

const exampleSentence = "I love cats and cats are so much fun!";
console.log("Original sentence:");
console.log(exampleSentence);

const dogsOnlySentence = exampleSentence.replaceAll("cats", "dogs");
console.log("Replacing all occurrences of cats with dogs:");
console.log(dogsOnlySentence);

const learningSentence = "I love learning!";
console.log("Original learning sentence:");
console.log(learningSentence);

const repeatedLove = "love ".repeat(3);
console.log(repeatedLove);

--fcc-editable-region--

--fcc-editable-region--
```
