---
id: 68e42b62d7093345c6975f22
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Abaixo da sua variável `originalString`, adicione um `console.log()` com a string `"Original string:"`. 

Então abaixo daquele `console.log()`, adicione outro `console.log()` com a variável `originalString`. 

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve registrar a string `"Original string:"` no console.

```js
assert.equal(getLogs()[0], "Original string:");
```

Você deve registrar a variável `originalString` no console. 

```js
const codeWithoutComments = __helpers.removeJSComments(code);
const loggingOriginalStr = codeWithoutComments.match(/console\.log\(\s*originalString\s*\)/g)

assert.isAtLeast(loggingOriginalStr.length, 1);
```

# --seed--

## --seed-contents--

```js
const originalString = "I love cats.";
--fcc-editable-region--

--fcc-editable-region--
```
