---
id: 68de40f53fcaf106f51cc20f
title: Construye un reproductor de audio y video HTML
challengeType: 25
dashedName: build-an-html-audio-and-video-player
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debería tener un elemento `h1` para el título principal de la página.
2. Debes tener dos elementos `section`.
3. Dentro del primer elemento `section`, debes tener un elemento `h2` para el título del video que se está reproduciendo.
4. Debajo del elemento `h2`, debes tener un elemento `video` con los atributos `controls` y `width`. El atributo `width` debe estar configurado en `640`.
5. Dentro del elemento `video`, debería tener un elemento `source` con un atributo `src` que apunte a un archivo de video y un atributo `type`.
   - Puedes usar `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4`.
6. Dentro del segundo elemento `section`, debes tener un elemento `h2` para el título de la canción que está sonando.
7. Debajo del elemento `h2`, debes tener un elemento `audio` con los atributos `controls` y `loop`, y un atributo `src` que apunte a un archivo de audio.
   - Puedes usar `https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3`.
   - O `https://cdn.freecodecamp.org/curriculum/js-music-player/we-are-going-to-make-it.mp3`.

# --hints--

Debería tener un elemento `h1` para el título principal de la página.

```js
const h1 = document.querySelector('h1');
assert.exists(h1);
```

Tu `h1` no debe estar vacío.

```js
const h1 = document.querySelector('h1');
assert.isNotEmpty(h1.textContent.trim());
```

Debes tener solo un elemento `h1`

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Debes tener dos elementos `section`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);

```

Dentro del primer elemento `section`, debes tener un elemento `h2` para el título del video que se está reproduciendo.

```js
const section = document.querySelector('section');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Debajo del elemento `h2`, debes tener un elemento `video`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
assert.exists(video);
```

El elemento `video` debe tener un atributo `controls`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const controlsAttr = video.controls;
assert.isTrue(controlsAttr);
```

El elemento `video` debe tener un atributo `width` establecido en `640`.

```js
const section = document.querySelector('section');
const video = section?.querySelector('h2 + video');
const widthAttr = video.getAttribute('width');
assert.equal(widthAttr, '640');
```

Dentro del elemento `video`, debes tener un elemento `source`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
assert.exists(source);
```

El elemento `source` debería tener un atributo `src` que apunte a un archivo de video.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const srcAttr = source.getAttribute('src');
assert.match(srcAttr, /\.(mp4|webm|ogg|avi|mov)$/i);
```

El elemento `source` debe tener un atributo `type`

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
assert.isNotEmpty(typeAttr);
```

El atributo `type` debe coincidir con el tipo de medio de la extensión de archivo del atributo `src`.

```js
const section = document.querySelector('section');
const source = section?.querySelector('h2 + video > source');
const typeAttr = source.getAttribute('type');
const srcAttr = source.getAttribute('src');
const srcFileExt = srcAttr.match(/[^\.]+$/g)[0];
const fileExtToMediaTypes = {
  mp4: ['video/mp4'],
  webm: ['video/webm'],
  ogg: ['video/ogg'],
  avi: ['video/avi', 'video/vnd.avi', 'video/msvideo', 'video/x-msvideo'],
  mov: ['video/quicktime']
};
const srcMediaTypes = fileExtToMediaTypes[srcFileExt];
assert.oneOf(typeAttr, srcMediaTypes);
```

Dentro del segundo elemento `section`, debes tener un elemento `h2` para el título de la canción que está sonando.

```js
const section = document.querySelector('section:nth-of-type(2)');
const h2 = section?.querySelector('h2');
assert.exists(h2);
```

Todos tus elementos `h2` deben contener algún texto.

```js
const h2s = document.querySelectorAll('h2');
assert.isNotEmpty(h2s);
h2s.forEach((h2) => assert.isNotEmpty(h2.textContent.trim()));
```

Debajo del elemento `h2`, debes tener un elemento `audio`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
assert.exists(audio);
```

El elemento `audio` debe tener un atributo `controls`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const controlsAttr = audio.controls;
assert.isTrue(controlsAttr);
```

El elemento `audio` debe tener un atributo `loop`.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const loopAttr = audio.loop;
assert.isTrue(loopAttr);
```

El elemento `audio` debe tener un atributo `src` que apunte a un archivo de audio.

```js
const section = document.querySelector('section:nth-of-type(2)');
const audio = section?.querySelector('h2 + audio');
const srcAttr = audio.getAttribute('src');
assert.match(srcAttr, /\.(mp3|wav|ogg|aac|flac)$/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>

</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8" />
  <meta name="viewport" content="width=device-width, initial-scale=1.0"/>
  <title>HTML Audio and Video Lab</title>
</head>
<body>
  <main>
    <h1>HTML Audio and Video Lab</h1>

    <section>
      <h2>What is the map method and how does it work?</h2>
      <video controls width="640">
        <source src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4" type="video/mp4" />
      </video>
    </section>

    <section>
      <h2>Sailing Away</h2>
      <audio src="https://cdn.freecodecamp.org/curriculum/js-music-player/sailing-away.mp3" controls loop>
      </audio>
    </section>

  </main>
</body>
</html>
```
