---
id: 685239f0cad43c7a909cbd98
title: 調試寵物領養頁面
challengeType: 25
dashedName: lab-debug-pet-adoption-page
---

# --description--

Sally 是一位寵物領養商店的老闆，她搭建了自己的第一個網頁，但存在一些問題。

你的任務是修復所有錯誤，以便 Sally 能夠繼續構建她的頁面。

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. Sally 想使用一張貓咪的圖像，但圖像顯示不正確。 你需要在 `img` 元素中修正以下內容：
   - 用圖像源的正確屬性替換 `href` 屬性。
   - 將 `att` 屬性替換爲代表圖像簡短描述性文本的正確屬性。
   - 移除 `</img>` 結束標籤，因爲 `img` 元素是空元素，沒有結束標籤。
2. Sally 希望使用一些鏈接將用戶引導到貓和狗的頁面。 但鏈接無法正常工作。 你需要在 `a` 元素中修正以下內容：
   - 將兩個 `src` 屬性替換爲用於指定 URL 的正確屬性。

# --hints--

你的 `img` 元素應該有一個 `src` 屬性，而不是 `href` 屬性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

你的 `img` 元素應該有一個 `alt` 屬性，而不是不存在的 `att` 屬性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

你的 `img` 元素不應該有 `</img>` 結束標籤。

```js
assert.notMatch(code, /<\/img>/);
```

你的文本爲 `Visit cats page` 的 `a` 元素應該有一個 `href` 屬性，而不是 `src` 屬性。

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

你的文本爲 `Visit dogs page` 的 `a` 元素應該有一個 `href` 屬性，而不是 `src` 屬性。

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
