---
id: 685239f0cad43c7a909cbd98
title: 偵錯寵物領養頁面
challengeType: 25
dashedName: lab-debug-pet-adoption-page
# Purposefully removed the demo type because I don't want any example to show since this is a debugging project.
---

# --description--

Sally，一位寵物領養商店老闆，已經建立了她的第一個網頁，但有一些問題。

你的工作是修正所有錯誤，讓 Sally 可以繼續建置她的頁面。 

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. Sally 想使用一些貓的影像，但它沒有正確顯示。你需要修正 `img` 元素中的以下問題：
   - 將 `href` 屬性取代為影像原始碼的正確屬性。
   - 將 `att` 屬性取代為正確表述影像簡短描述性文字的屬性。
   - 移除 `</img>` 結束標籤，因為 `img` 元素是空元素，沒有結束標籤。
2. Sally 想使用一些聯結將使用者導向狗和貓的頁面。但這些聯結無法正確運作。你需要修正 `a` 元素中的以下問題：
   - 將兩個 `src` 屬性取代為用來指定 URL 的正確屬性。

# --hints--

你的 `img` 元素應該有 `src` 屬性，而不是 `href` 屬性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("src"));
```

你的 `img` 元素應該有一個 `alt` 屬性，而不是不存在的 `att` 屬性。

```js
const imgEl = document.querySelector("img");
assert.isTrue(imgEl.hasAttribute("alt"));
```

你的 `img` 元素不應該有 `</img>` 結束標籤。

```js
assert.notMatch(code, /<\/img>/);
```

你的 `a` 元素中帶有 `Visit cats page` 文字，需要有 `href` 屬性，而不是 `src` 屬性。

```js
const anchors = document.querySelectorAll("a");
const catLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit cats page");
assert.isDefined(catLink);
assert.isTrue(catLink.hasAttribute("href"));
```

你的 `a` 元素中帶有文字 `Visit dogs page`，需要有一個 `href` 屬性，而不是 `src` 屬性。

```js
const anchors = document.querySelectorAll("a");
const dogLink = Array.from(anchors).find(a => a.textContent.trim() === "Visit dogs page");
assert.isDefined(dogLink);
assert.isTrue(dogLink.hasAttribute("href"));
```

# --seed--

## --seed-contents--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" att="Two tabby kittens sleeping together on a couch."></img>

<h2>Adopt a cat!</h2>
<a src="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a src="/dogs">Visit dogs page</a>
```

# --solutions--

```html
<h1>Welcome XYZ Pet Adoption!</h1>
<p>Consider adopting a pet today. We have cats, dogs, rabbits and more.</p>

<h2>See our cats!</h2>
<img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">

<h2>Adopt a cat!</h2>
<a href="/cats">Visit cats page</a>

<h2>Adopt a dog!</h2>
<a href="/dogs">Visit dogs page</a>
```
