---
id: 6650c9a94d6e13d14a043a69
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

開発者の「CamperBot」さんは、背景色ランダム変更プログラムを作成しようとしています。 しかし、次から次へと問題に直面している状況です。コードのデバッグを手伝ってあげましょう。

CamperBot さんは、既にプロジェクトのための HTML と CSS を追加しました。 ところが、スタイルもコンテンツもページに表示されず、困っています。

コンソールを開くと、次のメッセージが表示されます。

```js
SyntaxError: unknown: Unexpected token, expected "," (5:2)
```

構文エラー (syntax error) は、JavaScript エンジンが解釈できないものに遭遇したときに発生します。 このケースでは、CamperBot さんが作成した `darkColorsArr` という配列内に、構文エラーがあるようです。

`darkColorsArr` 配列内の構文エラーを修正してください。すると、コンテンツとスタイルがページに表示されるはずです。

# --hints--

`darkColorsArr` 配列内の構文エラーを解決する必要があります。 配列の要素はコンマで区切ることを忘れないでください。

```js
assert.isArray(darkColorsArr);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Build a random background color changer</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1>Random Background Color changer</h1>

    <main>
      <section class="bg-information-container">
        <p>Hex Code: <span id="bg-hex-code">#110815</span></p>
      </section>

      <button class="btn" id="btn">Change Background Color</button>
    </main>
    <script src="./script.js"></script>
  </body>
</html>
```

```css
*,
*::before,
*::after {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --yellow: #f1be32;
  --golden-yellow: #feac32;
  --dark-purple: #110815;
  --light-grey: #efefef;
}

body {
  background-color: var(--dark-purple);
  color: var(--light-grey);
  text-align: center;
}

.bg-information-container {
  margin: 15px 0 25px;
  font-size: 1.2rem;
}

.btn {
  cursor: pointer;
  padding: 10px;
  margin: 10px;
  color: var(--dark-purple);
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}

.btn:hover {
  background-image: linear-gradient(#ffcc4c, #f89808);
}

```

```js
--fcc-editable-region--
const darkColorsArr = [
  "#2C3E50",
  "#34495E",
  "#2C2C2C"
  "#616A6B"
  "#4A235A"
  "#2F4F4F",
  "#0E4B5A",
  "#36454F",
  "#2C3E50",
  "#800020",
];
--fcc-editable-region--
```
