---
id: 61329b210dac0b08047fd6ab
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

Continuing with the `meta` elements, a `viewport` definition tells the browser how to render the page. Including one betters visual accessibility on mobile, and improves _SEO_ (search engine optimization).

`viewport` 定義を追加し、その `content` 属性にはページの `width` と `initial-scale` の詳細を記述してください。

# --hints--

`head` 要素の中に新しい `meta` 要素を作成する必要があります。

```js
assert.equal(document.querySelectorAll('head > meta')?.length, 2);
```

`meta` タグに、値が `viewport` に設定されている `name` 属性を加える必要があります。

```js
assert.equal(document.querySelectorAll('head > meta[name="viewport"]')?.length, 1);
```

`meta` タグに、値が `width=device-width, initial-scale=1` に設定されている `content` 属性を加える必要があります。

```js
assert.equal(document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1.0"]')?.length || document.querySelectorAll('head > meta[content="width=device-width, initial-scale=1"]')?.length, 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
