---
id: 69b83e35f19ba26ba1fa517a
title: Costruisci un sequenziatore per semafori
challengeType: 26
dashedName: lab-traffic-light-sequencer
---

# --description--

In questo laboratorio simulerai cicli configurabili di semafori e registrerai le anomalie.

Lavorerai con oggetti di configurazione che descrivono le fasi di un semaforo. Ogni oggetto di configurazione ha le seguenti proprietĂ :

- `fault`: una flag booleana che attiva la terminazione anticipata quando Ă¨ `true`.
- `phases`: un array di oggetti fase.

Ogni oggetto fase dentro `phases` ha le seguenti proprietĂ :

- `color`: una stringa che rappresenta il colore della luce (`"green"`, `"yellow"` o `"red"`).
- `duration`: un intero positivo che indica la durata della fase in secondi.

Puoi fare riferimento agli oggetti `config1`, `config2`, `config3` e `config4` forniti come esempi di possibili oggetti di configurazione.

**Nota**: Non aggiungere ulteriori istruzioni `console.log()`, perchĂ© potrebbero far fallire i test.

**Obiettivo**: soddisfa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

# Storie utente

1. Dovresti avere una funzione chiamata `runSequence` con due parametri: `config` e `cycles`, dove `config` rappresenta un oggetto di configurazione e `cycles` rappresenta il numero massimo di volte in cui la sequenza puĂ˛ essere eseguita prima di fermarsi.

1. Dovresti implementare `runSequence(config, cycles)` usando un loop `for` o `while` per iterare su ogni fase per il numero di cicli indicato.

1. La funzione `runSequence` dovrebbe:
    - Registrare `No phases found` e restituire immediatamente se `config.phases` Ă¨ vuoto.
    - Registrare `Faulted phase!` e fermare la simulazione anticipatamente se `config.fault` Ă¨ `true`.
    - Registrare `Invalid phase detected` se `duration <= 0`.
    - Registrare `Switching to [color] for [duration] s` per ogni fase valida. Sostituisci `[color]` e `[duration]` con le proprietĂ  corrispondenti dell'oggetto fase.
    - Per esempio, `runSequence(config1, 1)` dovrebbe registrare:

  ```md
  Switching to green for 5 s
  Switching to yellow for 2 s
  Switching to red for 4 s
  ```

1. Dovresti avere una funzione chiamata `generateTimeline` con due parametri: `config` e `cycles`.

1. La funzione `generateTimeline` dovrebbe:
    - Registrare il tempo cumulativo trascorso dopo ogni fase per tutti i cicli in un array, sommando la `duration` di ogni fase al totale corrente mentre iteri.
    - Processare tutte le fasi con fault e invalide senza validazione, anche se `config.fault` Ă¨ `true` o `duration <= 0`.
    - Restituire l'array dei timestamp cumulativi.
    - Per esempio, `generateTimeline(config1, 1)` dovrebbe restituire l'array `[5, 7, 11]`.

# --before-each--

```js
const spy = __helpers.spyOn(console, "log");
```

# --hints--

Non dovresti modificare gli oggetti `config1`, `config2`, `config3` e `config4`.

```js
const _config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const _config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const _config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const _config4 = {
  fault: false,
  phases: []
};

assert.deepEqual(_config1, config1);
assert.deepEqual(_config2, config2);
assert.deepEqual(_config3, config3);
assert.deepEqual(_config4, config4);
```

Dovresti avere una funzione chiamata `runSequence`.

```js
assert.isFunction(runSequence);
```

La tua funzione `runSequence` dovrebbe accettare due parametri, `config` e `cycles`.

```js
const regex = __helpers.functionRegex('runSequence', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regex);
```

`runSequence(config1, 1)` dovrebbe registrare in ordine `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`.

```js
runSequence(config1, 1);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config1, 2)` dovrebbe registrare in ordine `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`, `Switching to green for 5 s`, `Switching to yellow for 2 s`, `Switching to red for 4 s`.

```js
runSequence(config1, 2);

assert.deepEqual(spy.calls, [
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"],
  ["Switching to green for 5 s"],
  ["Switching to yellow for 2 s"],
  ["Switching to red for 4 s"]
]);
```

`runSequence(config2, 1)` dovrebbe registrare in ordine `Switching to red for 3 s`, `Invalid phase detected`, `Switching to green for 6 s`.

```js
runSequence(config2, 1);

assert.deepEqual(spy.calls, [
  ["Switching to red for 3 s"],
  ["Invalid phase detected"],
  ["Switching to green for 6 s"]
]);
```

`runSequence(config3, 2)` dovrebbe registrare `Faulted phase!`.

```js
runSequence(config3, 2);

assert.deepEqual(spy.calls, [["Faulted phase!"]]);
```

`runSequence(config4, 5)` dovrebbe registrare `No phases found`.

```js
runSequence(config4, 5);

assert.deepEqual(spy.calls, [["No phases found"]]);
```

Dovresti avere una funzione chiamata `generateTimeline`.

```js
assert.isFunction(generateTimeline);
```

La tua funzione `generateTimeline` dovrebbe accettare due parametri, `config` e `cycles`.

```js
const regEx = __helpers.functionRegex('generateTimeline', ['config', 'cycles']);
assert.match(__helpers.removeJSComments(code), regEx);
```

`generateTimeline(config1, 1)` dovrebbe restituire `[5, 7, 11]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 1), [5, 7, 11]);
```

`generateTimeline(config1, 2)` dovrebbe restituire `[5, 7, 11, 16, 18, 22]`.

```js
assert.deepStrictEqual(generateTimeline(config1, 2), [5, 7, 11, 16, 18, 22]);
```

`generateTimeline(config2, 2)` dovrebbe restituire `[3, 1, 7, 10, 8, 14]`.

```js
assert.deepStrictEqual(generateTimeline(config2, 2), [3, 1, 7, 10, 8, 14]);
```

`generateTimeline(config3, 1)` dovrebbe restituire `[5, 7, 13]`.

```js
assert.deepStrictEqual(generateTimeline(config3, 1), [5, 7, 13]);
```

`generateTimeline(config4, 1)` dovrebbe restituire `[]`.

```js
assert.deepStrictEqual(generateTimeline(config4, 1), []);
```

# --seed--

## --seed-contents--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};
```

# --solutions--

```js
const config1 = {
  fault: false,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 4 }
  ]
};

const config2 = {
  fault: false,
  phases: [
    { color: "red", duration: 3 },
    { color: "yellow", duration: -2 },
    { color: "green", duration: 6 }
  ]
};

const config3 = {
  fault: true,
  phases: [
    { color: "green", duration: 5 },
    { color: "yellow", duration: 2 },
    { color: "red", duration: 6 }
  ]
};

const config4 = {
  fault: false,
  phases: []
};

function runSequence(config, cycles) {
  if (config.phases.length === 0) {
    console.log("No phases found");
    return;
  }
  if (config.fault === true) {
    console.log("Faulted phase!");
    return;
  }
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { color, duration } = config.phases[i];
      if (duration <= 0) {
        console.log("Invalid phase detected");
      }
      else {
        console.log(`Switching to ${color} for ${duration} s`);
      }
    }
  }
}

function generateTimeline(config, cycles) {
  let timeLine = [];
  let currTime = 0;
  for (let cycle = 1; cycle <= cycles; cycle++) {
    for (let i = 0; i < config.phases.length; i++) {
      let { duration } = config.phases[i];
      currTime += duration;
      timeLine.push(currTime);
    }
  }
  return timeLine;
}
```
