---
id: 66a37f37ef5823a313de8c26
title: Eine Mondumlaufbahn erstellen
challengeType: 25
dashedName: build-a-moon-orbit
demoType: onClick
---

# --description--

In diesem Labor erstellen Sie eine einfache Animation der Mondumlaufbahn um die Erde mit HTML und CSS. Die Erde befindet sich im Zentrum des Systems, und der Mond umkreist sie.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten alle Elemente zentrieren, indem Sie die Höhe auf `100%` des Viewports setzen und ein Flexbox-Layout im `body`-Element verwenden.
   
2. Sie sollten ein `div` mit der Klasse `space` haben.

3. Innerhalb des `.space`-Elements sollten sich zwei weitere `div`-Elemente mit den Klassen `earth` und `orbit` in dieser Reihenfolge befinden.

4. Im `.orbit`-Element sollte ein weiteres `div` mit der Klasse `moon` sein.
   
5. Das `div`-Element mit einer `class` von `space` sollte auf der Seite zentriert sein und eine Breite und Höhe von `200px` haben.

6. Das `div`-Element mit einer `class` von `space` sollte `relative`-Positionierung verwenden.

7. Das `.earth`-Element sollte `absolute`-Positionierung verwenden. Positionieren Sie das Zentrum des `.earth`-Elements auf der Hälfte seines übergeordneten Elements sowohl auf der vertikalen (oben) als auch auf der horizontalen (links) Achse. Verschieben Sie danach das `.earth`-Element um die Hälfte seiner eigenen Breite und Höhe zurück, um es innerhalb seines übergeordneten `.space`-Elements zu zentrieren.

8. Das `.earth`-Element sollte eine Breite und Höhe von `100px` haben.

9.  Das `.orbit`-Element sollte eine Breite und Höhe von `200px` haben.

10.  Das `.orbit`-Element sollte mit `absolute`-Positionierung positioniert sein. Seine untere rechte Ecke sollte sich mit einer `.space`-Eigenschaft, die es um `transform` auf beiden Achsen verschiebt, im Zentrum des `-50%`-Elements befinden.

11. Die Umlaufbahn des Mondes um die Erde sollte ein Kreis sein.

12. Das `.moon`-Element sollte `absolute`-Positionierung verwenden und eine Breite und Höhe von `30px` haben. Das `.moon`-Element sollte sich am oberen Rand des `.orbit`-Elements befinden und horizontal zentriert sein.

13. Sie sollten die horizontale Positionierung des `.moon`-Elements weiter anpassen, indem Sie das Element um die Hälfte seiner Breite nach links verschieben.

14. Der `.earth`-Selektor sollte eine Hintergrundfarbe und eine `border-radius` von `50%` haben.

15. Der `.moon`-Selektor sollte eine Hintergrundfarbe und eine `border-radius` von `50%` haben.

16. Sie sollten eine `@keyframes orbit`-Animation definieren, die das `.orbit`-Element 360 Grad um sein Zentrum dreht. Wenden Sie diese Animation auf das `.orbit`-Element mit einer Dauer von `5` Sekunden, einer linearen Timing-Funktion und unendlichen Wiederholungen an.

**Hinweis:** Vergewissern Sie sich, dass Sie Ihr Stylesheet in Ihrem HTML verlinken und Ihr CSS anwenden.

# --hints--

Der Inhalt des `body`-Elements sollte auf der Seite zentriert sein, indem die Höhe auf `100%` der Viewport-Höhe gesetzt wird, `display` auf `flex`, `justify-content` und `align-items` auf `center`.

```js

assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.display, 'flex');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.justifyContent, 'center');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('body')?.alignItems, 'center');
```

Sie sollten ein Haupt-`div` mit der Klasse `space` haben.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('class'), 'space');
```

Das `.space`-Element sollte eine Breite und Höhe von `200px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.height, '200px');
```

Das `.space`-Element sollte eine `position`-Eigenschaft von `relative` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.space')?.position, 'relative')
```

Sie sollten ein `div` mit der Klasse `earth` innerhalb des `div` mit einer `class` von `space` haben.

```js
assert.exists(document.querySelector('div.space div.earth'));
```

Das `.earth`-Element sollte das erste Kind des `.space`-Elements sein.

```js
assert.strictEqual(document.querySelector('.space > div')?.getAttribute('class'), 'earth');
```

Das `.earth`-Element sollte eine Breite und Höhe von `100px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.width, '100px');
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.height, '100px');
```

Das `.earth`-Element sollte `absolute`-Positionierung verwenden.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.position, 'absolute');
```

Das `.earth`-Element sollte eine `top`- und `left`-Position von `50%` haben.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.top, '50%') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.left, '50%');
```

Das `.earth`-Element sollte durch Setzen der `50%`-Eigenschaft auf `transform` um `translate(-50%, -50%)` nach hinten und oben verschoben werden.


```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.transform, 'translate(-50%, -50%)');
```

Sie sollten ein `div` mit der Klasse `orbit` innerhalb des `.space`-Elements haben.

```js
assert.exists(document.querySelector('div.space div.orbit'));
```

Das `.orbit`-Element sollte nach dem `.earth`-Element kommen.

```js
assert.isTrue(document.querySelector('.earth')?.nextElementSibling?.classList?.contains('orbit'));
```

Das `.orbit`-Element sollte eine Breite und Höhe von `200px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.width, '200px')
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.height, '200px');
```

Das `.orbit`-Element sollte `absolute`-Positionierung verwenden.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.position, 'absolute');
```

Das `.orbit`-Element sollte mit einer `.space`-Eigenschaft, die es um `transform` auf beiden Achsen verschiebt, im Zentrum des `-50%`-Elements positioniert sein.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.orbit')?.transform, 'translate(-50%, -50%)');
```

Sie sollten ein `div` mit der Klasse `moon` innerhalb des `.orbit`-Elements haben.

```js
assert.exists(document.querySelector('div.space div.orbit div.moon'));
```

Das `.moon`-Element sollte eine Breite und Höhe von `30px` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.width, '30px') 
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.height, '30px');
```

Das `.moon`-Element sollte `absolute`-Positionierung verwenden.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.position, 'absolute');
```

Die obere Kante des `.moon`-Elements sollte mit der oberen Kante seines enthaltenden Blocks ausgerichtet sein.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.top, '0px');
```

Sie sollten die linke Kante des `.moon`-Elements in die Mitte des enthaltenden Blocks verschieben (`50%` von der linken Seite).

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.left, '50%');
```

Das `.moon`-Element sollte horizontal angepasst werden, indem es um `-50%` entlang der x-Achse verschoben wird.

```js
const transformValue = new __helpers.CSSHelp(document).getStyle('.moon')?.transform;
const validTransforms = ['translateX(-50%)', 'translate(-50%)'];
assert.oneOf(transformValue, validTransforms);
```

Der `.earth`-Selektor sollte eine Hintergrundfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.earth')?.backgroundColor);
```

Der `.earth`-Selektor sollte eine `border-radius` von `50%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.earth')?.borderRadius, '50%');
```

Der `.moon`-Selektor sollte eine Hintergrundfarbe haben.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.moon')?.backgroundColor);
```

Der `.moon`-Selektor sollte eine `border-radius` von `50%` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.moon')?.borderRadius, '50%');
```

Sie sollten eine `@keyframes`-Regel haben.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 1);
```

Ihre neue `@keyframes`-Regel sollte `orbit` heißen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'orbit');
```

Ihre `@keyframes orbit`-Regel sollte einen `0%`-Selektor haben.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules; 
assert.strictEqual(rules?.[0]?.keyText, '0%');
```

Ihre `@keyframes orbit`-Regel sollte einen `100%`-Selektor haben.

```js 
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

Ihr `0%`-Selektor sollte eine `transform`-Eigenschaft mit dem Wert `rotate(0deg) translate(-50%, -50%)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[0]?.style?.transform, 'rotate(0deg) translate(-50%, -50%)');
```

Ihr `100%`-Selektor sollte nach Ihrem `0%`-Selektor kommen.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.keyText, '100%')
```

Ihr `100%`-Selektor sollte eine `transform`-Eigenschaft mit dem Wert `rotate(360deg) translate(-50%, -50%)` haben.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.cssRules?.[1]?.style?.transform, 'rotate(360deg) translate(-50%, -50%)')
```

Ihr `.orbit`-Selektor sollte eine `animation`-Eigenschaft mit dem Wert `orbit 5s linear infinite` haben.

```js

const orbitStyles = new __helpers.CSSHelp(document).getStyle('.orbit');
assert.strictEqual(orbitStyles?.animationName, "orbit");
assert.strictEqual(orbitStyles?.animationDuration, "5s");
assert.strictEqual(orbitStyles?.animationTimingFunction, "linear");
assert.strictEqual(orbitStyles?.animationIterationCount, "infinite");

```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Moon Orbit</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Moon Orbiting Earth</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <div class="space">
        <div class="earth"></div>
        <div class="orbit">
            <div class="moon"></div>
        </div>
    </div>
</body>

</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #000;
}

.space {
    position: relative;
    width: 200px;
    height: 200px;
}

.earth {
    width: 100px;
    height: 100px;
    background-color: blue;
    border-radius: 50%;
    position: absolute;
    top: 50%;
    left: 50%;
    transform: translate(-50%, -50%);
    box-shadow: 0 0 20px rgba(0, 0, 255, 0.5);
}

.orbit {
    width: 200px;
    height: 200px;
    position: absolute;
    transform: translate(-50%, -50%);
    animation: orbit 5s linear infinite;
}

.moon {
    width: 30px;
    height: 30px;
    background-color: gray;
    border-radius: 50%;
    position: absolute;
    top: 0;
    left: 50%;
    transform: translateX(-50%);
    box-shadow: 0 0 10px rgb(255, 255, 255);
}


@keyframes orbit {
    0% {
        transform: rotate(0deg) translate(-50%, -50%);
    }

    100% {
        transform: rotate(360deg) translate(-50%, -50%);
    }
}
```

