---
id: 6989eb65324d97775c4965fe
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 너비 우선 탐색(BFS) 방식을 사용해 소괄호의 모든 유효한 조합을 생성하는 함수를 구현할 거예요. 예를 들어, 두 쌍의 소괄호로 만들 수 있는 유효한 조합은 `(())`와 `()()`입니다.

먼저 `genParentheses`라는 이름의 함수를 만들고, 매개변수는 하나인 `pairs`로 하세요. 지금은 함수에서 빈 배열을 반환하도록 하세요.

# --hints--

`genParentheses`라는 이름의 함수를 정의해야 합니다.

```js
assert.isFunction(genParentheses);
```

`genParentheses` 함수는 `pairs`라는 하나의 매개변수를 가져야 합니다.

```js
const genParenthesesParams = __helpers.getFunctionParams(genParentheses.toString());
assert.equal(genParenthesesParams[0].name, 'pairs')
```

함수는 빈 배열을 반환해야 합니다.

```js
assert.isArray(genParentheses(1));
assert.lengthOf(genParentheses(1), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
