---
id: 6989f207306332c06b0add39
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

너비 우선 탐색 방식을 위해, 소괄호 조합을 만들면서 다양한 상태를 추적할 큐를 사용하세요. 각 상태는 세 요소를 포함하는 배열로 나타냅니다:

- 현재 만들어지고 있는 문자열
- 지금까지 사용한 여는 소괄호 개수
- 지금까지 사용한 닫는 소괄호 개수

`queue`라는 변수를 선언하고, 빈 문자열과 소괄호가 0개 사용된 시작 상태를 나타내는 배열 `['', 0, 0]`을 포함하는 배열로 초기화하세요.

# --hints--

`queue` 함수 안에 `genParentheses`라는 변수를 선언해야 합니다.

```js
assert.match(genParentheses.toString(), /queue/);
```

`queue`를 배열 `['', 0, 0]`를 포함하는 배열로 초기화해야 합니다.

```js
assert.match(genParentheses.toString(), /queue\s*=\s*\[\s*\['',\s*0,\s*0\]\s*\]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  let result = [];
  return result;
}
```
