---
id: 6989f3632f8b8bdf168aa321
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

`while` 반복문 안에서 `queue.shift()`을 사용해 큐에서 첫 번째 요소를 제거하고 가져오세요. 이렇게 하면 BFS의 특징인 선입선출(FIFO) 동작을 구현할 수 있습니다.

이 배열을 `current`, `opensUsed`, `closesUsed` 세 변수로 구조 분해하세요.

# --hints--

큐에서 첫 번째 요소를 가져오려면 `queue.shift()`을 사용하세요.

```js
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

결과를 `current`, `opensUsed`, `closesUsed` 세 변수로 구조 분해해야 합니다.

```js
assert.match(genParentheses.toString(), /opensUsed/);
assert.match(genParentheses.toString(), /closesUsed/);
assert.match(genParentheses.toString(), /current/);
assert.match(genParentheses.toString(), /queue\.shift\s*\(\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return result;
}
```
