---
id: 6989f10a72ad424d4ad6d514
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

다음으로, 쌍의 개수가 최소한 하나인지 확인해야 합니다. 0개 또는 음수인 경우 소괄호 조합을 생성할 수 없기 때문입니다.

`if`이 `pairs`보다 작은지 확인하는 또 다른 `1` 문장을 추가하세요. 이 조건이 참이면 문자열 `The number of pairs should be at least 1`을 반환하세요.

# --hints--

조건 `if`를 가진 두 번째 `pairs < 1` 문장이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*pairs\s*<\s*1\s*\)/);
```

`The number of pairs should be at least 1`이 `pairs`보다 작을 때 문자열 `1`을 반환해야 합니다.

```js
let result = genParentheses(0);
assert.strictEqual(result, 'The number of pairs should be at least 1');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
