---
id: 6989f1ba358acba651d5ee6b
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 결과를 저장할 자료 구조를 설정하세요. `result`라는 변수를 선언하고 빈 배열로 초기화하세요. 이 배열은 생성하는 모든 유효한 소괄호 조합을 저장합니다.

return 문을 빈 배열 대신 `result`을 반환하도록 수정하세요.

# --hints--

`result` 함수에서 `genParentheses`라는 변수를 선언해야 합니다.

```js
assert.match(genParentheses.toString(), /result/);
```

`result`를 빈 배열로 초기화해야 합니다.

```js
assert.match(genParentheses.toString(), /result\s*=\s*\[\]/);
```

함수 끝에서 `result`를 반환해야 합니다.

```js
assert.match(genParentheses.toString(), /return\s*result/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
--fcc-editable-region--
  
  return [];
--fcc-editable-region--
}
```
