---
id: 6989f2e35e6941e13e445156
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 메인 BFS 루프를 구현하세요. `while`이 비어 있지 않은 동안, 즉 `queue`가 `queue.length`보다 큰 동안 계속되는 `0` 루프를 만드세요.

루프 안에서 `queue`를 콘솔에 출력하세요.

# --hints--

조건 `while`을 가진 `queue.length > 0` 루프를 만들어야 합니다.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
```

`queue` 루프 안에서 `while`를 콘솔에 출력해야 합니다.

```js
assert.match(genParentheses.toString(), /while\s*\(\s*queue\.length\s*>\s*0\s*\)/);
assert.match(genParentheses.toString(), /console\.log\s*\(\s*queue\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }
  let queue = [['', 0, 0]];
  let result = [];
--fcc-editable-region--
  
--fcc-editable-region--
    
  return result;
}
```
