---
id: 6989f5bf7f1a1d0fc8ebd1e8
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 닫는 소괄호를 추가하는 부분을 처리해야 합니다. 핵심 규칙은 균형을 유지하는 경우에만 닫는 소괄호를 추가할 수 있다는 것입니다. 즉, 닫는 소괄호보다 여는 소괄호가 더 많이 사용되어야 합니다.

여전히 `else` 블록 안에서, 두 번째 `if` 문장을 추가하여 `closesUsed`가 `opensUsed`보다 작은지 확인하세요.

이 조건이 참이면, 닫는 소괄호를 추가한 후 상태를 나타내는 새로운 배열을 `queue`에 `[current + ')', opensUsed, closesUsed + 1]`로 추가하세요.

# --hints--

`if` 블록 안에 조건 `closesUsed < opensUsed`을 가진 두 번째 `else` 문장이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*closesUsed\s*<\s*opensUsed\s*\)/);
```

조건이 참일 때 `[current + ')', opensUsed, closesUsed + 1]`를 큐에 추가해야 합니다.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\)['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
  return result;
}

console.log(genParentheses(1));
```
