---
id: 6989f711e57008f3a23482f2
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

함수가 이제 완성되었습니다. `genParentheses(2)` 대신 `genParentheses(1)`을 기록하여 다시 테스트하세요.

# --hints--

`genParentheses(2)`을 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*genParentheses\s*\(\s*2\s*\)\s*\)/);
```

코드에 `console.log(genParentheses(1))`이 없어야 합니다.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\.log\s*\(\s*genParentheses\s*\(\s*1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();

    if (current.length === 2 * pairs) {
      result.push(current);
    } else {
      if (opensUsed < pairs) {
        queue.push([current + '(', opensUsed + 1, closesUsed]);
      }
      if (closesUsed < opensUsed) {
        queue.push([current + ')', opensUsed, closesUsed + 1]);
      }
    }
  }
  return result;
}

--fcc-editable-region--
console.log(genParentheses(1));
--fcc-editable-region--
```
