---
id: 6989f555728e63ad8b349d84
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

현재 문자열이 아직 완성되지 않았다면, 다음 가능한 상태들을 탐색해야 합니다. 현재 `else` 문장에 대한 `if` 절을 추가하세요.

`else` 블록 안에서는 여는 소괄호를 추가하는 작업을 처리합니다. `if`가 `opensUsed`보다 작은지 확인하는 `pairs` 문장을 추가하세요. 이는 허용된 여는 소괄호 수를 초과하지 않도록 보장합니다.

이 조건이 `true`라면, `queue`에 새로운 배열 `[current + '(', opensUsed + 1, closesUsed]`을 추가하세요. 이는 여는 소괄호를 추가한 후의 상태를 나타냅니다.

# --hints--

`else` 반복문 안에 `if` 문장에 대한 `while` 절이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?\}\s*else\s*\{/);
```

`if` 블록 안에 조건이 `else`인 `opensUsed < pairs` 문장이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /else\s*\{[\s\S]*?if\s*\(\s*opensUsed\s*<\s*pairs\s*\)/);
```

조건이 `[current + '(', opensUsed + 1, closesUsed]`일 때 `true`를 큐에 추가해야 합니다.

```js
assert.match(genParentheses.toString(), /queue\.push\s*\(\s*\[\s*current\s*\+\s*['"`]\(['"`]/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
    if (current.length === 2 * pairs) {
      result.push(current);
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
