---
id: 6989f46b6e9366b0e6ad32fe
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

BFS 로직을 계속 진행하기 전에 `genParentheses(1)`를 호출하고 결과를 콘솔에 기록하세요.

# --hints--

`genParentheses(1)`를 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*genParentheses\s*\(\s*1\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
  }

  return result;
}

--fcc-editable-region--

--fcc-editable-region--
```
