---
id: 6989f48c3f0292c84f8d8938
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이제 완전한 소괄호 조합을 만들었을 때를 식별해야 합니다. 완전한 조합은 쌍의 수의 두 배에 해당하는 길이를 가지는데, 각 쌍이 여는 소괄호 하나와 닫는 소괄호 하나를 포함하기 때문입니다.

`while` 반복문 안에 `if` 문장을 추가하여 `current.length === 2 * pairs`인지 확인하세요. 이 조건이 참일 때 `current`을 `result` 배열에 추가하세요.

# --hints--

`if` 반복문 안에 조건 `current.length === 2 * pairs`를 가진 `while` 문장이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)/);
```

새로운 `current` 문장 안에 `result`을 `if` 배열에 추가해야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*current\.length\s*===?\s*2\s*\*\s*pairs\s*\)\s*\{[\s\S]*?result\.push\s*\(\s*current\s*\)/);
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
  if (!Number.isInteger(pairs)) {
    return 'The number of pairs should be an integer';
  }
  if (pairs < 1) {
    return 'The number of pairs should be at least 1';
  }

  let queue = [['', 0, 0]];
  let result = [];

  while (queue.length > 0) {
    console.log(queue);
    let [current, opensUsed, closesUsed] = queue.shift();
--fcc-editable-region--
    
--fcc-editable-region--
  }
  return result;
}

console.log(genParentheses(1));
```
