---
id: 6989ef885db3a3488a13b426
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

핵심 알고리즘을 구현하기 전에 입력을 검증해야 합니다. `pairs` 매개변수는 생성할 소괄호 쌍의 개수를 나타내므로 정수여야 합니다.

함수 시작 부분에 `if` 문을 추가하여 `pairs`가 정수가 아닌지 확인하세요. 이를 위해 `Number.isInteger()` 메서드를 사용하세요.

조건이 참이면 문자열 `The number of pairs should be an integer`를 반환하세요.

# --hints--

`!Number.isInteger(pairs)` 함수 안에 조건이 `genParentheses`인 if 문이 있어야 합니다.

```js
assert.match(genParentheses.toString(), /if\s*\(\s*!\s*Number\.isInteger\s*\(\s*pairs\s*\)\s*\)/);
```

입력이 정수가 아닐 때 문자열을 반환해야 합니다.

```js
let result = genParentheses("i");
assert.isString(result);
```

입력이 문자열일 때 문자열 `The number of pairs should be an integer`를 반환해야 합니다.

```js
let result = genParentheses("i");
assert.strictEqual(result, 'The number of pairs should be an integer');
```

# --seed--

## --seed-contents--

```js
function genParentheses(pairs) {
--fcc-editable-region--
  
--fcc-editable-region--
  return [];
}
```
