---
id: 669e2f60e83c011754f711f9
title: Створіть сторінку туристичної агенції
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

Fulfill the user stories below and get all the tests to pass to complete the lab.

**Історія користувача:**

1. Ви повинні мати оголошення `DOCTYPE`.
1. Ви повинні мати елемент `html` з `lang` зі значенням `en`.
1. Ви повинні мати елемент `head`, який містить порожній елемент `meta` з `charset` зі значенням `utf-8` та `title` із текстом `Travel Agency Page`.
1. Ви повинні мати елемент `h1`, щоб представити пункти призначення вашої подорожі.
1. Ви повинні мати абзац під елементом `h1`, який представляє можливості подорожей.
1. Ви повинні мати елемент `h2` із текстом `Packages`.
1. Ви повинні мати елемент `p`, який коротко представляє різні пакунки.
1. Ви повинні мати невпорядкований елемент списку з двома пунктами списку. Два пункти списку повинні мати текст `Group Travels` та `Private Tours` відповідно. Текст кожного пункту списку повинен бути закритий елементом посилання.
1. Ви повинні мати елемент `h2` із текстом `Top Itineraries`.
1. Ви повинні мати принаймні три елементи `figure`, кожен з яких містить елемент посиланння та елемент `figcaption`.
1. Три елементи посиланння повинні мати елемент `img` з відповідним атрибутом `alt` та атрибутом `src`, встановленим на дійсне зображення як їх вміст. Ви можете використовувати `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` та `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`, якщо хочете.
1. Усі п’ять елементів посиланння повинні мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn` та атрибут `target` зі значенням `_blank`.

# --hints--

На сторінці вашої туристичної агенції має бути оголошення `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Ви повинні мати елемент `html` з `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `head` в межах елемента `html`.

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

Ви повинні мати елемент `meta` в межах елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

Ви повинні мати елемент `title` в межах елемента `head`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

Елемент `title` повинен містити назву туристичного агентства.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

Ви повинні мати елемент `body` в межах елемента `html`.

```js
assert(code.match(/<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i));
```

Ви повинні мати елемент `h1`, щоб представити пункти призначення вашої подорожі.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0)
```

Ви повинні мати тільки один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1)
```

You should have a `p` element below the `h1` element.

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'))
```

Your first `p` element should introduce the travel opportunities.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Перший елемент `h2` повинен містити текст `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Packages');
```

You should have a `p` element below your first `h2` element.

```js
assert.strictEqual(document.querySelector('h2')?.nextElementSibling, document.querySelectorAll('p')?.[1])
```

Your second `p` element should introduce briefly the various packages.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim())
```

Під другим елементом `p` має бути невпорядкований елемент списку.

```js
assert.strictEqual(document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName, 'UL');
```

У невпорядкованому списку мають бути два елементи.

```js
assert.equal(document.querySelectorAll('li').length,  2);
```

Обидва елементи списку повинні містити елемент посиланння.

```js
const listItems = $('li');
assert(!!listItems.length);
for (let e of listItems) {
  assert(e.children[0].tagName === 'A' && e.children.length === 1);
}
```

Елемент посиланння першого елемента списку має обгортати текст `Group Travels`.

```js
assert.equal(document.querySelectorAll('a')[0]?.innerText, 'Group Travels');
```

Елемент посиланння другого елемента списку має обгортати текст `Private Tours`.

```js
assert.equal(document.querySelectorAll('a')[1]?.innerText, 'Private Tours');
```

Ви повинні мати елемент `h2` після невпорядкованого списку.

```js
assert($('h2')[1]?.previousElementSibling.tagName === 'UL');
```

Другий елемент `h2` повинен містити текст `Top Itineraries`.

```js
assert.equal($('h2')[1]?.innerText, 'Top Itineraries');
```

Ви повинні мати принаймні три елементи `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3)
```

Кожен елемент `figure` повинен містити елемент посилання як перший дочірній елемент.

```js
const figures = $('figure');
assert(!!figures.length);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName,'A');
}
```

Кожен елемент `figure` повинен містити елемент `figcaption` як другий дочірній елемент.

```js
const figures = $('figure');
assert(!!figures.length);
for (let figure of figures) {
  assert.equal(figure.children[1].tagName, 'FIGCAPTION');
}
```

Кожен з елементів `a`, які є дітьми елементів `figure`, повинен містити зображення.

```js
const anchors = $('figure > a');
assert(!!anchors.length);
for (let anchor of anchors) {
  assert.equal(anchor.children[0].tagName, 'IMG');
}
```

Кожен елемент `img` повинен мати дійсний атрибут `src`.

```js
const images = $('img');
assert(!!images.length);
for (let image of images) {
    assert.isAbove(image.src.length, 0);
}
```

Кожен елемент `img` повинен мати атрибут `alt` із відповідним значенням.

```js
const images = $('img');
assert(!!images.length);
for (let image of images) {
    assert.isAbove(image.alt.length, 0);
}
```

Кожен елемент `a` повинен мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn`.

```js
const anchors = $('a');
assert(!!anchors.length);
for (let anchor of anchors) {
    assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Кожен елемент `a` повинен мати атрибут `target` зі значенням `_blank`.

```js
const anchors = $('a');
assert(!!anchors.length);
for (let anchor of anchors) {
    assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
    <head>
        <meta charset="utf-8">
        <title>Travel Agency Page</title>
    </head>
    <body>
        <h1>Discover Italy</h1>
        <p>Art, folklore, food, nature, and more. Choose among our wide selection of guided tours and excursions, and live an unforgettable experience exploring Italy.</p>
        <h2>Packages</h2>
        <p>We offer an extensive range of holiday solutions to accommodate the needs of all our clients. From daily excursions in the most beautiful cities, to thorough tours of hidden villages and medieval towns to discover Italy's lesser-known sides.</p>
        <ul>
            <li><a href="https://www.freecodecamp.org/learn" target="_blank">Group Travels</a></li>
            <li><a href="https://www.freecodecamp.org/learn" target="_blank">Private Tours</a></li>
        </ul>
        <h2>Top Itineraries</h2>
        <figure>
            <a href="https://www.freecodecamp.org/learn" target="_blank">
                <img src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg" alt="colosseum view" >                
            </a>
            <figcaption>Rome and Center Italy</figcaption>
        </figure>
        <figure>
            <a href="https://www.freecodecamp.org/learn" target="_blank">
                <img src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg" alt="dolomites mountain view" >                
            </a>
            <figcaption>Nature and National Parks</figcaption>            
        </figure>
        <figure>
            <a href="https://www.freecodecamp.org/learn" target="_blank">
                <img src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg" alt="sea view" >                
            </a>
            <figcaption>South Italy and Islands</figcaption>
        </figure>                
    </body>
</html>
```
