---
id: 669e2f60e83c011754f711f9
title: Створіть сторінку туристичної агенції
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

Виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати оголошення `DOCTYPE`.
1. Ви повинні мати елемент `html` з `lang` зі значенням `en`.
1. Ви повинні мати елемент `head`, який містить порожній елемент `meta` з атрибутом `charset` зі значенням `utf-8` та `title` з текстом `Travel Agency Page`.
1. Ви повинні мати тег `meta` в елементі `head`, який містить короткий опис вебсайту для SEO.
1. Ви повинні мати елемент `h1`, де представлено напрямок подорожі.
1. Під елементом `h1` має бути абзац, де написано короткий вступ.
1. Ви повинні мати елемент `h2` з текстом `Packages`.
1. Ви повинні мати елемент `p`, де коротко описано різні тури.
1. Ви повинні мати невпорядкований список з двома елементами. Два елементи списку відповідно повинні мати текст `Group Travels` та `Private Tours`. Текст цих елементів має бути закритий посиланням.
1. Ви повинні мати елемент `h2` з текстом `Top Itineraries`.
1. Ви повинні мати принаймні три елементи `figure`, кожен з яких містить посилання та елемент `figcaption`.
1. Три посилання повинні мати елемент `img` з відповідним атрибутом `alt` і атрибутом `src` зі значенням дійсного зображення. Якщо хочете, можете використати `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`, `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg` та `https://cdn.freecodecamp.org/curriculum/labs/sea.jpg`.
1. Усі п’ять посилань повинні мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn` і атрибут `target` зі значенням `_blank`.

# --hints--

На сторінці туристичної агенції має бути оголошення `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

Ви повинні мати елемент `html` з `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/i);
```

В межах елемента `html` має бути елемент `head`.

```js
assert.match(
  code,
  /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i
);
```

В межах елемента `head` має бути два елементи `meta`.

```js
assert.match(code, /<\s*head\s*>[\s\S]*(<\s*meta[\s\S]*>[\s\S]*){2,}<\/\s*head\s*>/i);
```

Один елемент `meta` повинен мати атрибут `name` зі значенням `description` і непорожній атрибут `content`.

 ```js
 assert.match(code,
  /<\s*meta\s+[^>]*name\s*=\s*["']description["'][^>]*content\s*=\s*["'][^"']+?["'][^>]*>/i,
);
```

Один елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

В межах елемента `head` має бути елемент `title`.

```js
assert.match(
  code,
  /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i
);
```

Елемент `title` має містити назву туристичної агенції.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

В межах елемента `html` має бути елемент `body`.

```js
assert.match(
  code,
  /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i
);
```

Ви повинні мати елемент `h1`, де представлено напрямок подорожі.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Ви повинні мати лише один елемент `h1`.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Під елементом `h1` має бути елемент `p`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

У першому елементі `p` має бути короткий вступ.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

Перший елемент `h2` повинен містити текст `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Packages');
```

Під першим елементом `h2` має бути елемент `p`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

Другий елемент `p` має коротко описувати різні тури.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

Під другим елементом `p` має бути невпорядкований список.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

У невпорядкованому списку мають бути два елементи.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

Обидва елементи списку мають містити посилання.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

Посилання першого елемента списку має обгортати текст `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText, 'Group Travels');
```

Посилання другого елемента списку має обгортати текст `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText, 'Private Tours');
```

Після невпорядкованого списку має бути елемент `h2`.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

Другий елемент `h2` повинен містити текст `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Top Itineraries');
```

Ви повинні мати принаймні три елементи `figure`.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

Першим дочірнім елементом всіх елементів `figure` має бути посилання.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

Другим дочірнім елементом всіх елементів `figure` має бути елемент `figcaption`.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

Всі елементи `figcaption` мають містити текст.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

Всі елементи `a`, які є дочірніми елементами елементів `figure`, повинні містити зображення.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

Всі елементи `img` повинні мати дійсний атрибут `src`.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

Всі елементи `img` повинні мати атрибут `alt` з відповідним значенням.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

Всі елементи `a` повинні мати атрибут `href` зі значенням `https://www.freecodecamp.org/learn`. Не забудьте про посилання в елементах списку.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

Всі елементи `a` повинні мати атрибут `target` зі значенням `_blank`. Не забудьте про посилання в елементах списку.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
