---
id: 5f33071498eb2472b87ddee4
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Projekt lernen Sie die Grundlagen von CSS (Cascading Style Sheets), indem Sie ein Café-Menü erstellen. CSS ist die Sprache, mit der ein HTML-Dokument gestaltet wird. Es beschreibt, wie HTML-Elemente auf dem Bildschirm angezeigt werden sollen.

Wie Sie in den letzten Schritten der Cat Photo App gelernt haben, gibt es eine Grundstruktur, die benötigt wird, um Ihre Webseite zu erstellen. Jedes HTML-Dokument sollte eine `DOCTYPE` Deklaration und ein `html` Element haben. Die `DOCTYPE` teilt dem Browser mit, welche Version von HTML im Dokument verwendet wird. Und das `html` Element stellt das Wurzelelement dar, das alle anderen Elemente enthält.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Fügen Sie den `<!DOCTYPE html>` Tag hinzu und ein `html` Element mit einem `lang` Attribut von `en`.

# --hints--

Sie sollten die `<!DOCTYPE html>` Deklaration haben.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Sie sollten einen öffnenden `<html lang="en">` mit der Sprache auf Englisch eingestellt haben.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Sie sollten einen schließenden `</html>` Tag haben.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
