---
id: 5f332b23c2045fb843337579
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Um den Besuchern mitzuteilen, dass das Café im Jahr 2020 gegründet wurde, fügen Sie ein `p`-Element unter dem `h1`-Element mit dem Text `Est. 2020` hinzu.

# --hints--

Sie sollten ein öffnendes `<p>`-Tag haben.

```js
assert(code.match(/<p>/i));
```

Sie sollten ein schließendes `</p>`-Tag haben.

```js
assert(code.match(/<\/p>/i));
```

Sie sollten Ihr bestehendes `h1`-Element nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

Ihr `p`-Element sollte unter Ihrem `h1`-Element stehen.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

Ihr `p`-Element sollte den Text `Est. 2020` enthalten.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

