---
id: 5f33294a6af5e9188dbdb8f3
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Es ist Zeit, etwas Menüinhalt hinzuzufügen. Fügen Sie ein `main`-Element innerhalb des vorhandenen `body`-Elements ein. Es wird später Preisinformationen über Kaffee und Desserts enthalten, die das Café anbietet.

# --hints--

Ihr Code sollte ein öffnendes `<main>`-Tag enthalten.

```js
assert(code.match(/<main>/i));
```

Ihr Code sollte ein schließendes `</main>`-Tag enthalten.

```js
assert(code.match(/<\/main>/i));
```

Sie sollten Ihr `body`-Element nicht ändern. Achten Sie darauf, das schließende Tag nicht versehentlich zu löschen.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Ihr `main`-Tag sollte sich innerhalb Ihres `body`-Tags befinden.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
