---
id: 5f3326b143638ee1a09ff1e3
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Um sich darauf vorzubereiten, echten Inhalt zu erstellen, fügen Sie ein `body`-Element unterhalb des `head`-Elements hinzu.

# --hints--

Sie sollten ein öffnendes `<body>`-Tag haben.

```js
assert(code.match(/<body>/i));
```

Sie sollten ein schließendes `</body>`-Tag haben.

```js
assert(code.match(/<\/body>/i));
```

Sie sollten Ihr `head`-Element nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

Ihr `body`-Element sollte nach Ihrem `head`-Element kommen.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

