---
id: 5f331e55dfab7a896e53c3a1
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Das `title` ist eines von mehreren Elementen, die zusätzliche Informationen bereitstellen, die auf der Webseite nicht sichtbar sind, aber für Suchmaschinen oder die Darstellung der Seite nützlich sind.

Verschachteln Sie innerhalb des `head`-Elements ein `meta`-Element mit einem Attribut namens `charset`, das auf den Wert `utf-8` gesetzt ist, um dem Browser mitzuteilen, wie Zeichen für die Seite kodiert werden sollen.

Denken Sie daran, dass `meta`-Elemente leere Elemente sind, die keinen End-Tag benötigen.

# --hints--

Sie sollten einen `meta`-Tag haben.

```js
assert.match(code, /<meta.*?\/?>/is);
```

Das `meta`-Element ist ein leeres Element, es sollte keinen End-Tag `</meta>` haben.

```js
assert.notMatch(code, /<\/meta>/i);
```

Ihr `meta`-Tag sollte ein `charset`-Attribut haben.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Ihr `charset`-Attribut sollte den Wert `utf-8` haben.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
