---
id: 5f3313e74582ad9d063e3a38
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie ein `head`-Element innerhalb des `html`-Elements hinzu, damit Sie ein `title`-Element hinzufügen können. Der Text des `title`-Elements sollte `Cafe Menu` sein.

# --hints--

Sie sollten ein öffnendes `<head>`-Tag haben.

```js
assert(code.match(/<head>/i));
```

Sie sollten ein schließendes `</head>`-Tag haben.

```js
assert(code.match(/<head>/i));
```

Sie sollten ein öffnendes `<title>`-Tag haben.

```js
assert(code.match(/<title>/i));
```

Sie sollten ein schließendes `</title>`-Tag haben.

```js
assert(code.match(/<\/title>/i));
```

Ihr `title`-Element sollte in Ihrem `head`-Element verschachtelt sein.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

Ihr `title`-Element sollte den Text `Cafe Menu` haben. Überprüfen Sie gegebenenfalls Ihre Rechtschreibung.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
