---
id: 5f332a88dc25a0fd25c7687a
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Der Name des Cafés ist `CAMPER CAFE`. Fügen Sie ein `h1`-Element innerhalb Ihres `main`-Elements hinzu. Geben Sie ihm den Namen des Cafés in Großbuchstaben, damit er hervorsticht.

# --hints--

Sie sollten ein öffnendes `<h1>`-Tag haben.

```js
assert(code.match(/<h1>/i));
```

Sie sollten ein schließendes `</h1>`-Tag haben.

```js
assert(code.match(/<\/h1>/i));
```

Sie sollten Ihr bestehendes `main`-Element nicht ändern.

```js
assert.lengthOf(document.querySelectorAll('main'), 1);
```

Ihr `h1`-Element sollte in Ihrem `main`-Element verschachtelt sein.

```js
assert.equal(document.querySelector('h1')?.parentElement?.tagName, "MAIN");
```

Ihr `h1`-Element sollte den Text `CAMPER CAFE` enthalten.

```js
assert(code.match(/<h1>CAMPER CAFE<\/h1>/));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
--fcc-editable-region--
    <main>
    </main>
--fcc-editable-region--
  </body>
</html>
```

