---
id: 5f3477ae8466a9a3d2cc953c
title: Schritt 15
challengeType: 0
dashedName: step-15
---

# --description--

Da Sie nun das CSS in der `styles.css`-Datei haben, entfernen Sie bitte das `style`-Element und seinen gesamten Inhalt. Sobald es entfernt ist, wird der zuvor zentrierte Text wieder nach links verschoben.

# --hints--

In Ihrem Code sollten keine `style`-Tags mehr vorhanden sein.

```js
assert(!code.match(/style/i));
```

In Ihrer HTML-Datei sollten keine CSS-Selektoren enthalten sein.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

