---
id: 5f344fad8bf01691e71a30eb
title: Schritt 10
challengeType: 0
dashedName: step-10
---

# --description--

Bis jetzt hatten Sie nur begrenzte Kontrolle über die Darstellung und das Aussehen Ihres Inhalts. Um das zu ändern, fügen Sie ein `style`-Element innerhalb des `head`-Elements hinzu.

# --hints--

Ihr Code sollte einen öffnenden `<style>`-Tag enthalten.

```js
assert(code.match(/<style\s*>/i));
```

Ihr Code sollte einen schließenden `</style>`-Tag enthalten.

```js
assert(code.match(/<\/style\s*>/));
```

Ihr `style`-Element sollte in Ihrem `head`-Element verschachtelt sein.

```js
assert(code.match(/<head\s*>[\w\W\s]*<style\s*>[\w\W\s]*<\/style\s*>[\w\W\s]*<\/head\s*>/i))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
--fcc-editable-region--
</html>
```
