---
id: 5f3477aefa51bfc29327200b
title: Schritt 14
challengeType: 0
dashedName: step-14
---

# --description--

Sie haben drei Elemente gestylt, indem Sie CSS innerhalb der `style` Tags geschrieben haben. Das funktioniert, aber da noch viele weitere Styles folgen werden, ist es am besten, alle Styles in einer separaten Datei zu speichern und mit dieser zu verknüpfen.

Für Sie wurde bereits eine separate `styles.css` Datei erstellt. Sie können mit den Tabs oben im Editor zwischen den Dateien wechseln.

Beginnen Sie damit, die Styles, die Sie erstellt haben, in die `styles.css` Datei umzuschreiben. Achten Sie darauf, die öffnenden und schließenden `style` Tags auszuschließen.

# --hints--

Ihre `styles.css` Datei sollte den `h1, h2, p` Typselektor enthalten.

```js
assert(code.slice(code.indexOf('</html>')).replace(/[\s\n]*/g, "").match(__helpers.concatRegex(__helpers.permutateRegex(['h1', 'h2', 'p'], { elementsSeparator: ',' }), /{/)));
```
 
Ihr Selektor sollte die `text-align` Eigenschaft auf `center` setzen.

```js
assert(code.slice(code.indexOf('</html>')).match(/text-align:\s*center;?/));
```

Ihr Code sollte keine `<style>` Tags enthalten.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
