---
id: 5f344fbc22624a2976425065
title: Schritt 9
challengeType: 0
dashedName: step-9
---

# --description--

Erstellen Sie ein `h2`-Element im `section`-Element und geben Sie ihm den Text `Coffee`.

# --hints--

Sie sollten ein öffnendes `<h2>`-Tag haben.

```js
assert(code.match(/<h2\s*>/i));
```

Sie sollten ein schließendes `</h2>`-Tag haben.

```js
assert(code.match(/<\/h2\s*>/i));
```

Sie sollten Ihr bestehendes `section`-Element nicht ändern. Stellen Sie sicher, dass Sie das schließende Tag nicht gelöscht haben.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Ihr `h2`-Element sollte sich innerhalb Ihres `section`-Elements befinden.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Ihr `h2`-Element sollte den Text `Coffee` haben.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
