---
id: 5f3477cb303c5cb61b43aa9b
title: Schritt 18
challengeType: 0
dashedName: step-18
---

# --description--

Der Text ist wieder zentriert, sodass die Verknüpfung mit der CSS-Datei funktioniert. Fügen Sie der Datei einen weiteren Stil hinzu, der die `background-color`-Eigenschaft für das `brown`-Element in `body` ändert.

# --hints--

Sie sollten einen `body`-Selektor verwenden.

```js
const hasBody = new __helpers.CSSHelp(document).getStyle('body');
assert(hasBody);
```

Sie sollten die `background-color`-Eigenschaft auf `brown` setzen.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-color'] === 'brown');
assert(hasBackground);
```

Ihr `body`-Element sollte einen `brown`-Hintergrund haben.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-color');
assert(bodyBackground === 'brown');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--
h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

