---
id: 5f35e5c4321f818cdc4bed30
title: Schritt 29
challengeType: 0
dashedName: step-29
---

# --description--

Da jetzt alles gut aussieht, ist es Zeit, einige Menüpunkte hinzuzufügen.

Fügen Sie unter der `article`-Überschrift ein leeres `Coffee`-Element hinzu. Es soll den Geschmack und den Preis jedes Kaffees enthalten, den Sie derzeit anbieten.


# --hints--

Sie sollten ein öffnendes `<article>`-Tag haben.

```js
assert.match(code,/<article>/i);
```

Sie sollten ein schließendes `</article>`-Tag haben.

```js
assert(code.match(/<\/article>/i));
```

Sie sollten das bestehende `h2`-Element nicht ändern.

```js
assert.lengthOf(document.querySelectorAll('h2'),1);
```

Ihr `article`-Element sollte nach Ihrem `h2`-Element kommen.

```js
const article = document.querySelector('article');
assert.equal(article?.previousElementSibling?.tagName, 'H2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
--fcc-editable-region--
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

