---
id: 5f356ed60785e1f3e9850b6e
title: Schritt 24
challengeType: 0
dashedName: step-24
---

# --description--

Jetzt ist leicht zu erkennen, dass der Text innerhalb des `#menu`-Elements zentriert ist. Derzeit ist die Breite des `#menu`-Elements in Pixeln (`px`) angegeben.

Ändern Sie den Wert der `width`-Eigenschaft auf `80%`, um die Breite auf 80 % des übergeordneten Elements (`body`) festzulegen.

# --hints--

Sie sollten die `width`-Eigenschaft auf `80%` setzen.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '80%');
assert(hasWidth);
```

Sie sollten keine `width`-Eigenschaft mit dem Wert `300px` haben.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(!hasWidth);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--
#menu {
  width: 300px;
  background-color: burlywood;
}
--fcc-editable-region--
```

