---
id: 5f3c866dd0d0275f01d4d847
title: Schritt 39
challengeType: 0
dashedName: step-39
---

# --description--

Nun, das hat nicht funktioniert. Wenn Sie die `p`-Elemente als `inline-block` stylen und sie auf separate Zeilen setzen, entsteht rechts neben dem ersten `p`-Element ein zusätzlicher Abstand, wodurch das zweite Element in die nächste Zeile rutscht.

Eine Möglichkeit, das zu beheben, besteht darin, die Breite jedes `p`-Elements etwas kleiner als `50%` zu machen. Ändern Sie also den Wert von `width` auf `49%` für jede Klasse, um zu sehen, was passiert.

# --hints--

Sie sollten die `width`-Eigenschaft in Ihrem `49%`-Selektor auf `.flavor` setzen.

```js
const flavorWidth = new __helpers.CSSHelp(document).getStyle('.flavor')?.getPropertyValue('width');
assert(flavorWidth === '49%');
```

Sie sollten die `width`-Eigenschaft in Ihrem `49%`-Selektor auf `.price` setzen.

```js
const priceWidth = new __helpers.CSSHelp(document).getStyle('.price')?.getPropertyValue('width');
assert(priceWidth === '49%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
          <article class="item">
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.item p {
  display: inline-block;
}

--fcc-editable-region--
.flavor {
  text-align: left;
  width: 50%;
}

.price {
  text-align: right;
  width: 50%;
}
--fcc-editable-region--
```

