---
id: 6140cbeec34e970dfe75e710
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

أنشئ محدِّدًا لعنصر `.wheel`. ابدأ بتعيين `border` إلى `2px solid black`، و`border-radius` إلى `50%`، و`margin-left` إلى `50px`.

# --hints--

يجب أن يكون لديك محدِّد `.wheel`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel'));
```

يجب أن يحتوي محدِّد `.wheel` على خاصية `border` مضبوطة على `2px solid black`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.border === '2px solid black');
```

يجب أن يحتوي محدِّد `.wheel` على خاصية `border-radius` مضبوطة على `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.borderRadius === '50%');
```

يجب أن يحتوي محدِّد `.wheel` على خاصية `margin-left` مضبوطة على `50px`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.wheel')?.marginLeft === '50px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
