---
id: 6140c7e645d8e905819f1dd4
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ بالقالب الجاهز القياسي. أضف إعلان `DOCTYPE`، وعنصر `html` مع تعيين اللغة إلى الإنجليزية، وعنصري `head` و`body`.

أضف عنصر `meta` الخاص بك مع `charset` الصحيح، وعنصر `title`، وعنصر `link` لملف `./styles.css`.

عيّن `title` إلى `Ferris Wheel`.

# --hints--

يجب أن تحتوي شفرتك على إعلان `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

يجب أن تضع مسافة بعد إعلان `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

يجب أن تحدد نوع المستند ليكون `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

يجب أن يحتوي عنصر `html` على وسم بداية مع خاصية `lang` بقيمة `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert.match(code, /<\/html\s*>/);
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

يجب أن يكون لديك وسم بداية `head`.

```js
assert.match(code, /<head\s*>/i);
```

يجب أن يكون لديك وسم إغلاق `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

يجب أن يكون لديك وسم بداية `body`.

```js
assert.match(code, /<body\s*>/i);
```

يجب أن يكون لديك وسم إغلاق `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

يجب أن يكون عنصر `body` هو العنصر الأول بعد عنصر `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

يجب أن يكون عنصرا `head` و`body` أخوين.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

يجب أن يكون عنصر `head` هو العنصر الأول داخل عنصر `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

يجب أن تحتوي شفرتك على عنصر `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

يجب أن يحتوي عنصر `meta` على خاصية `charset` بقيمة `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

يجب أن يكون عنصر `meta` داخل عنصر `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

يجب أن تحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

يجب أن يكون لمشروعك عنوان `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

تذكر، أن حالة الأحرف والتهجئة مهمة للعنوان.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

يجب أن تحتوي شفرتك على عنصر `link`.

```js
assert.match(code, /<link/);
```

يجب أن يكون لديك عنصر `link` واحد.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

يجب أن يكون عنصر `link` داخل عنصر `head`.

```js
assert.exists(document.querySelector('head > link'));
```

يجب أن يحتوي عنصر `link` على خاصية `rel` بقيمة `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

يجب أن يحتوي عنصر `link` على خاصية `href` بقيمة `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
