---
id: 6140d0069049f5139d78da40
title: الخطوة 8
challengeType: 0
dashedName: step-8
---

# --description--

تُستخدم خاصية `transform-origin` لتحديد النقطة التي يُطبق حولها تحويل CSS. على سبيل المثال، عند تنفيذ `rotate` (وهو ما ستفعله لاحقًا في هذا المشروع)، تحدد `transform-origin` النقطة التي يدور حولها العنصر.

أعطِ محدِّد `.line` خاصية `transform-origin` بقيمة `0% 0%`. هذا سيُزاح نقطة الأصل بمقدار `0%` من اليسار و`0%` من الأعلى، مما يضبطها على الزاوية العلوية اليسرى للعنصر.

# --hints--

يجب أن يحتوي محدِّد `.line` على خاصية `transform-origin` مضبوطة على `0% 0%`.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert(transformOrigin === '0% 0%' || transformOrigin === '0% 0% 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
