---
id: 6140cfc08ca9c5128c3e6478
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

عيّن خاصية `position` لمحدِّد `.line` إلى `absolute`، وعيّن خاصية `left` إلى `50%`، وخصيصة `top` إلى `50%`.

# --hints--

يجب أن يحتوي محدِّد `.line` على خاصية `position` مضبوطة على `absolute`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.position === 'absolute');
```

يجب أن يحتوي محدِّد `.line` على خاصية `left` مضبوطة على `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.left === '50%');
```

يجب أن يحتوي محدِّد `.line` على خاصية `top` مضبوطة على `50%`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.line')?.top === '50%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
}
--fcc-editable-region--
```
