---
id: 587d7fb7367417b2b2512c0c
title: Usa model.findOne() para devolver un único documento coincidente de tu base de datos
challengeType: 2
forumTopicId: 301545
dashedName: use-model-findone-to-return-a-single-matching-document-from-your-database
---

# --description--

`Model.findOne()` se comporta como `Model.find()`, pero solo devuelve un documento (no un arreglo), incluso si hay varios elementos. Es especialmente útil a la hora de buscar por propiedades que has declarado como únicas.

# --instructions--

Modifica la función `findOneByFood` para encontrar solo una persona que tenga cierta comida en los favoritos de la persona, usando `Model.findOne() -> Person`. Usa el argumento de la función `food` como clave de búsqueda.

# --hints--

Encontrar un elemento por Id debe ser exitoso

```js
  const response = await fetch(code + '/_api/find-one-by-food', {
    method: 'POST',
    headers: {
      'Content-Type': 'application/json'
    },
    body: JSON.stringify({
      name: 'Gary',
      age: 46,
      favoriteFoods: ['chicken salad']
    })
  });
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'Gary', 'item.name is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['chicken salad'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

