---
id: 587d7fb7367417b2b2512c0d
title: Usa model.findById() para buscar en tu base de datos por _id
challengeType: 2
forumTopicId: 301544
dashedName: use-model-findbyid-to-search-your-database-by-id
---

# --description--

Al guardar un documento, MongoDB añade automáticamente el campo `_id`, y lo establece como una clave alfanumérica única. Buscar por `_id` es una operación extremadamente frecuente, así que Mongoose proporciona un método dedicado para ello.

# --instructions--

Modifica el `findPersonById` para encontrar a la única persona que tenga un `_id` dado, usando `Model.findById() -> Person`. Usa el argumento de la función `personId` como clave de búsqueda.

# --hints--

Encontrar un elemento por Id debe ser exitoso

```js
  const response = await fetch(code + '/_api/find-by-id');
  if (!response.ok) {
    throw new Error(await response.text());
  }
  const data = await response.json();
  assert.equal(data.name, 'test', 'item.name is not what expected');
  assert.equal(data.age, 0, 'item.age is not what expected');
  assert.deepEqual(
    data.favoriteFoods,
    ['none'],
    'item.favoriteFoods is not what expected'
  );
  assert.equal(data.__v, 0, 'The item should be not previously edited');
```

