---
id: 587d7daf367417b2b2512b7e
title: Зрозуміти властивість конструктора
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

Існує спеціальна властивість `constructor`, розташована на екземплярах об’єктів `duck` та `beagle`, які були створені в попередніх завданнях:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

Обидва ці виклики `console.log` виведуть у консоль `true`.

Зверніть увагу, що властивість `constructor` є референсом на конструктор функції, яка створила цей екземпляр. Перевага властивості `constructor` полягає в тому, що можна перевірити цю властивість, щоб дізнатися, який це тип об’єкта. Ось приклад, як це можна використати:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**Примітка:** Оскільки властивість `constructor` може бути перезаписана (про це йтиметься в наступних двох завданнях), зазвичай краще використовувати метод `instanceof` для перевірки типу об’єкта.

# --instructions--

Напишіть функцію `joinDogFraternity`, яка приймає параметр `candidate` і, використовуючи властивість `constructor`, повертає `true`, якщо кандидат є `Dog`, інакше повертає `false`.

# --hints--

`joinDogFraternity` має бути визначена як функція.

```js
assert(typeof joinDogFraternity === 'function');
```

`joinDogFraternity` має повертати `true`, якщо `candidate` є екземпляром `Dog`.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

`joinDogFraternity` має використовувати властивість `constructor`.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
