---
id: 587d7db2367417b2b2512b89
title: Використання міксина для додавання загальної поведінки між несхожими об’єктами
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

Як ви вже бачили, поведінка передається через успадкування. Однак є випадки, коли успадкування не є найкращим рішенням. Успадкування не підходить для несхожих об’єктів, як `Bird` і `Airplane`. Вони обидва можуть літати, але `Bird` не є типом `Airplane` і навпаки.

Для несхожих об’єктів краще використовувати <dfn>міксини</dfn>. Міксин дозволяє іншим об’єктам використовувати набір функцій.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

`flyMixin` приймає будь-який об’єкт і надає йому метод `fly`.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

Тут `bird` і `plane` передаються в `flyMixin`, який потім призначає функцію `fly` кожному об’єкту. Тепер `bird` і `plane` можуть обидва літати:

```js
bird.fly();
plane.fly();
```

У консолі двічі відобразиться рядок `Flying, wooosh!`, по одному разу для кожного виклику `.fly()`.

Зверніть увагу, як міксин дозволяє повторно використовувати один і той самий метод `fly` для несхожих об’єктів `bird` і `plane`.

# --instructions--

Створіть міксин з іменем `glideMixin`, який визначає метод `glide`. Потім використайте `glideMixin`, щоб надати і `bird`, і `boat` здатність ковзати.

# --hints--

Ваш код має оголосити змінну `glideMixin`, яка є функцією.

```js
assert(typeof glideMixin === 'function');
```

Ваш код має використати `glideMixin` на об’єкті `bird`, щоб надати йому метод `glide`.

```js
assert(typeof bird.glide === 'function');
```

Ваш код має використати `glideMixin` на об’єкті `boat`, щоб надати йому метод `glide`.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
