---
id: 587d7db2367417b2b2512b8c
title: Використання IIFE для створення модуля
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

Одразу викликаний функціональний вираз (IIFE) часто використовують, щоб згрупувати пов’язану функціональність в один об’єкт або <dfn>модуль</dfn>. Наприклад, у попередньому завданні було визначено два міксини:

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

Ми можемо згрупувати ці міксини в модуль так:

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

Зверніть увагу, що у вас є одразу викликаний функціональний вираз (IIFE), який повертає об’єкт `motionModule`. Цей повернутий об’єкт містить усі поведінки міксинів як пропи об’єкта. Перевага шаблону модуля в тому, що всі поведінки руху можна упакувати в один об’єкт, який потім можна використовувати в інших частинах вашого коду. Ось приклад його використання:

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

Створіть модуль з іменем `funModule`, щоб обгорнути два міксини `isCuteMixin` і `singMixin`. `funModule` має повертати об’єкт.

# --hints--

`funModule` має бути визначений і повертати об’єкт.

```js
assert(typeof funModule === 'object');
```

`funModule.isCuteMixin` має звертатися до функції.

```js
assert(typeof funModule.isCuteMixin === 'function');
```

`funModule.singMixin` має звертатися до функції.

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
