---
id: 587d7dad367417b2b2512b78
title: Використання конструктора для створення об’єктів
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

Ось конструктор `Bird` з попереднього завдання:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**ПРИМІТКА:** `this` всередині конструктора завжди посилається на об’єкт, що створюється.

Зверніть увагу, що оператор `new` використовується при виклику конструктора. Це повідомляє JavaScript створити новий екземпляр `Bird` з ім’ям `blueBird`. Без оператора `new` `this` всередині конструктора не вказував би на новостворений об’єкт, що призвело б до несподіваних результатів. Тепер `blueBird` має всі властивості, визначені всередині конструктора `Bird`:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

Як і будь-який інший об’єкт, його властивості можна отримувати та змінювати:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

Використайте конструктор `Dog` з останнього уроку, щоб створити новий екземпляр `Dog`, призначивши його змінній `hound`.

# --hints--

`hound` має бути створений за допомогою конструктора `Dog`.

```js
assert(hound instanceof Dog);
```

Ваш код має використовувати оператор `new` для створення екземпляра `Dog`.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
