---
id: 587d7db2367417b2b2512b8b
title: Зрозуміти одразу викликаний функціональний вираз (IIFE)
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

Загальний шаблон у JavaScript — виконати функцію одразу після її оголошення:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

Це анонімний функціональний вираз, який виконується одразу і негайно виводить `Chirp, chirp!`.

Зверніть увагу, що функція не має імені і не зберігається у змінній. Дві круглі дужки () в кінці функціонального виразу змушують її одразу виконатися або викликатися. Цей шаблон відомий як <dfn>одразу викликаний функціональний вираз</dfn> або <dfn>IIFE</dfn>.

# --instructions--

Перепишіть функцію `makeNest` і видаліть її виклик, щоб замість цього це був анонімний одразу викликаний функціональний вираз (IIFE).

# --hints--

Функція має бути анонімною.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

Ваша функція має мати круглі дужки в кінці виразу, щоб викликати її одразу.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
