---
id: 587d7db0367417b2b2512b81
title: "Зрозуміти, звідки береться прототип об’єкта"
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

Так само, як люди успадковують гени від своїх батьків, об’єкт успадковує свій `prototype` безпосередньо від функції конструктора, яка його створила. Наприклад, тут конструктор `Bird` створює об’єкт `duck`:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

`duck` успадковує свій `prototype` від функції конструктора `Bird`. Ви можете показати цей зв’язок за допомогою методу `isPrototypeOf`:

```js
Bird.prototype.isPrototypeOf(duck);
```

Це поверне `true`.

# --instructions--

Використайте `isPrototypeOf`, щоб перевірити `prototype` у `beagle`.

# --hints--

Ви повинні показати, що `Dog.prototype` є `prototype` для `beagle`

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
