---
id: 587d7dac367417b2b2512b74
title: Використовуйте точкову нотацію для доступу до властивостей об’єкта
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

У попередньому завданні було створено об’єкт із різними властивостями. Тепер ви побачите, як отримати доступ до значень цих властивостей. Ось приклад:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

Точкова нотація використовується на імені об’єкта, `duck`, за яким слідує ім’я властивості, `name`, щоб отримати значення `Aflac`.

# --instructions--

Виведіть обидві властивості об’єкта `dog` у консоль.

# --hints--

Ваш код має використовувати `console.log` для виведення значення властивості `name` об’єкта `dog`.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

Ваш код має використовувати `console.log` для виведення значення властивості `numLegs` об’єкта `dog`.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
