---
id: 587d7db1367417b2b2512b85
title: Встановіть прототип дочірнього елемента як екземпляр батьківського елемента
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

У попередньому завданні ви побачили перший крок для наслідування поведінки від супертипу (або батьківського елемента) `Animal`: створення нового екземпляра `Animal`.

Це завдання охоплює наступний крок: встановити `prototype` підтипу (або дочірнього елемента) — у цьому випадку `Bird` — як екземпляр `Animal`.

```js
Bird.prototype = Object.create(Animal.prototype);
```

Пам’ятайте, що `prototype` — це як «рецепт» для створення об’єкта. У певному сенсі, рецепт для `Bird` тепер включає всі ключові «інгредієнти» з `Animal`.

```js
let duck = new Bird("Donald");
duck.eat();
```

`duck` наслідує всі властивості `Animal`, включно з методом `eat`.

# --instructions--

Змініть код так, щоб екземпляри `Dog` наслідували від `Animal`.

# --hints--

`Dog.prototype` має бути екземпляром `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
