---
id: 587d7dad367417b2b2512b75
title: Створити метод в об’єкті
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

Об’єкти можуть мати особливий тип властивості, який називається <dfn>метод</dfn>.

Методи — це властивості, які є функціями. Це додає об’єкту різну поведінку. Ось `duck` приклад з методом:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

У прикладі додано метод `sayName`, який є функцією, що повертає речення з ім’ям `duck`. Зверніть увагу, що метод отримав доступ до властивості `name` в операторі return, використовуючи `duck.name`. Наступне завдання розгляне інший спосіб зробити це.

# --instructions--

Використовуючи об’єкт `dog`, додайте йому метод з назвою `sayLegs`. Метод має повертати речення `This dog has 4 legs.`

# --hints--

`dog.sayLegs()` має бути функцією.

```js
assert(typeof dog.sayLegs === 'function');
```

`dog.sayLegs()` має повертати заданий рядок — зверніть увагу на пунктуацію та пробіли.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
