---
id: 587d7dae367417b2b2512b7b
title: розуміння власних пропертів
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

У наведеному прикладі конструктор `Bird` визначає дві проперти: `name` та `numLegs`:

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

`name` і `numLegs` називаються <dfn>власними пропертями</dfn>, тому що вони визначені безпосередньо на екземплярі об’єкта. Це означає, що `duck` і `canary` мають кожен свою окрему копію цих пропертів. Насправді кожен екземпляр `Bird` матиме свою копію цих пропертів. Наступний код додає всі власні проперти `duck` до масиву `ownProps`:

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

У консолі буде відображено значення `["name", "numLegs"]`.

# --instructions--

Додайте власні проперти `canary` до масиву `ownProps`.

# --hints--

`ownProps` має містити значення `numLegs` і `name`.

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

Ви маєте розв’язати це завдання без використання вбудованого методу `Object.keys()`.

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

Ви маєте розв’язати це завдання без жорсткого кодування масиву `ownProps`.

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
