---
id: 587d7dae367417b2b2512b7a
title: Перевірка конструктора об’єкта за допомогою instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

Кожного разу, коли конструктор створює новий об’єкт, цей об’єкт вважається <dfn>екземпляром</dfn> свого конструктора. JavaScript надає зручний спосіб перевірити це за допомогою оператора `instanceof`. `instanceof` дозволяє порівняти об’єкт із конструктором, повертаючи `true` або `false` залежно від того, чи був цей об’єкт створений за допомогою конструктора. Ось приклад:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

Цей `instanceof` метод поверне `true`.

Якщо об’єкт створено без використання конструктора, `instanceof` підтвердить, що він не є екземпляром цього конструктора:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

Цей `instanceof` метод поверне `false`.

# --instructions--

Створіть новий екземпляр конструктора `House`, назвіть його `myHouse` і передайте кількість спалень. Потім використайте `instanceof`, щоб перевірити, що це екземпляр `House`.

# --hints--

`myHouse` має мати атрибут `numBedrooms`, встановлений у число.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

Ви маєте перевірити, що `myHouse` є екземпляром `House` за допомогою оператора `instanceof`.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
