---
id: 587d7db0367417b2b2512b82
title: Зрозуміти ланцюжок прототипів
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

Усі об’єкти в JavaScript (з кількома винятками) мають `prototype`. Також `prototype` об’єкта сам є об’єктом.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

Оскільки `prototype` є об’єктом, `prototype` може мати власний `prototype`! У цьому випадку `prototype` `Bird.prototype` є `Object.prototype`:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

Навіщо це корисно? Ви, можливо, пам’ятаєте метод `hasOwnProperty` з попереднього завдання:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

Метод `hasOwnProperty` визначений у `Object.prototype`, до якого можна отримати доступ через `Bird.prototype`, а потім через `duck`. Це приклад ланцюжка `prototype`. У цьому ланцюжку `prototype` `Bird` є `supertype` для `duck`, тоді як `duck` є `subtype`. `Object` є `supertype` для обох `Bird` і `duck`. `Object` є `supertype` для всіх об’єктів у JavaScript. Тому будь-який об’єкт може використовувати метод `hasOwnProperty`.

# --instructions--

Змініть код, щоб показати правильний ланцюжок прототипів.

# --hints--

Ваш код має показати, що `Object.prototype` є прототипом `Dog.prototype`

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
