---
id: 587d7db1367417b2b2512b86
title: Скинути успадковану властивість конструктора
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

Коли об’єкт успадковує свій `prototype` від іншого об’єкта, він також успадковує властивість конструктора суперкласу.

Ось приклад:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

Але `duck` і всі екземпляри `Bird` повинні показувати, що вони були створені `Bird`, а не `Animal`. Для цього ви можете вручну встановити властивість конструктора `Bird` на об’єкт `Bird`:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

Виправте код так, щоб `duck.constructor` і `beagle.constructor` повертали свої відповідні конструктори.

# --hints--

`Bird.prototype` має бути екземпляром `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

`duck.constructor` має повертати `Bird`.

```js
assert(duck.constructor === Bird);
```

`Dog.prototype` має бути екземпляром `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

`beagle.constructor` має повертати `Dog`.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
