---
id: 587d7dac367417b2b2512b73
title: Створіть базовий об’єкт JavaScript
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

Подумайте про речі, які люди бачать щодня, як-от автомобілі, магазини та птахи. Це всі <dfn>об’єкти</dfn>: відчутні речі, які люди можуть спостерігати та з якими можуть взаємодіяти.

Які властивості мають ці об’єкти? Автомобіль має колеса. Магазини продають товари. Птахи мають крила.

Ці властивості, або <dfn>властивості</dfn>, визначають, з чого складається об’єкт. Зверніть увагу, що подібні об’єкти мають однакові властивості, але можуть мати різні значення цих властивостей. Наприклад, усі автомобілі мають колеса, але не всі автомобілі мають однакову кількість коліс.

Об’єкти в JavaScript використовують для моделювання реальних об’єктів, надаючи їм властивості та поведінку, як у їхніх реальних аналогів. Ось приклад, який використовує ці поняття для створення `duck` об’єкта:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

Цей `duck` об’єкт має дві пари властивість/значення: `name` зі значенням `Aflac` і `numLegs` зі значенням 2.

# --instructions--

Створіть `dog` об’єкт із властивостями `name` і `numLegs`, і встановіть їх відповідно в рядок і число.

# --hints--

`dog` має бути об’єктом.

```js
assert(typeof dog === 'object');
```

`dog` має мати властивість `name`, встановлену в рядок.

```js
assert(typeof dog.name === 'string');
```

`dog` має мати властивість `numLegs`, встановлену в число.

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
