---
id: 587d7dad367417b2b2512b77
title: Визначення функції конструктора
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>Конструктори</dfn> — це функції, які створюють нові об’єкти. Вони визначають властивості та поведінку, які належатимуть новому об’єкту. Уявіть їх як креслення для створення нових об’єктів.

Ось приклад конструктора:

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

Цей конструктор визначає об’єкт `Bird` з властивостями `name`, `color` та `numLegs`, значення яких відповідно Albert, синій і 2. Конструктори дотримуються кількох конвенцій:

<ul><li>Конструктори визначаються з іменем, що починається з великої літери, щоб відрізняти їх від інших функцій, які не є <code>конструкторами</code>.</li><li>Конструктори використовують ключове слово <code>this</code> для встановлення властивостей об’єкта, який вони створять. Всередині конструктора <code>this</code> посилається на новий об’єкт, який він створить.</li><li>Конструктори визначають властивості та поведінку замість того, щоб повертати значення, як це можуть робити інші функції.</li></ul>

# --instructions--

Створіть конструктор `Dog` з властивостями `name`, `color` та `numLegs`, які відповідно встановлені як рядок, рядок і число.

# --hints--

`Dog` має мати властивість `name`, встановлену як рядок.

```js
assert(typeof new Dog().name === 'string');
```

`Dog` має мати властивість `color`, встановлену як рядок.

```js
assert(typeof new Dog().color === 'string');
```

`Dog` має мати властивість `numLegs`, встановлену як число.

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
