---
id: 587d7db1367417b2b2512b87
title: додати методи після успадкування
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

Конструкторна функція, яка успадковує свій `prototype` об’єкт від конструктора суперкласу, все одно може мати власні методи на додаток до успадкованих.

Наприклад, `Bird` — це конструктор, який успадковує свій `prototype` від `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

Окрім того, що успадковується від `Animal`, ви хочете додати поведінку, унікальну для об’єктів `Bird`. Тут `Bird` отримає `fly()` функцію. Функції додаються до `Bird's` `prototype` так само, як і до будь-якої конструкторної функції:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

Тепер екземпляри `Bird` матимуть як `eat()`, так і `fly()` методи:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

`duck.eat()` виведе рядок `nom nom nom` у консоль, а `duck.fly()` виведе рядок `I'm flying!`.

# --instructions--

Додайте весь необхідний код, щоб об’єкт `Dog` успадковував від `Animal`, а конструктор `Dog` `prototype` був встановлений на `Dog`. Потім додайте метод `bark()` до об’єкта `Dog`, щоб `beagle` міг як `eat()`, так і `bark()`. Метод `bark()` має виводити `Woof!` у консоль.

# --hints--

`Animal` не повинен реагувати на метод `bark()`.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

`Dog` має успадкувати метод `eat()` від `Animal`.

```js
assert(typeof Dog.prototype.eat == 'function');
```

Прототип `Dog` має мати метод `bark()`.

```js
assert('bark' in Dog.prototype);
```

`beagle` має бути `instanceof` `Animal`.

```js
assert(beagle instanceof Animal);
```

Конструктор для `beagle` має бути встановлений на `Dog`.

```js
assert(beagle.constructor === Dog);
```

`beagle.eat()` має виводити рядок `nom nom nom`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

`beagle.bark()` має виводити рядок `Woof!`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
