---
id: 587d7daf367417b2b2512b80
title: Пам’ятайте встановлювати властивість constructor при зміні прототипу
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

Існує один важливий побічний ефект від ручного встановлення прототипу на новий об’єкт. Це стирає властивість `constructor`! Цю властивість можна використовувати, щоб перевірити, яка функція-конструктор створила екземпляр, але оскільки властивість була перезаписана, вона тепер дає хибні результати:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

Послідовно ці вирази оціняться як `false`, `true` і `true`.

Щоб виправити це, щоразу, коли прототип вручну встановлюють на новий об’єкт, пам’ятайте визначати властивість `constructor`:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

Визначте властивість `constructor` на `Dog` `prototype`.

# --hints--

`Dog.prototype` має встановити властивість `constructor`.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
