---
id: 587d7db0367417b2b2512b84
title: Успадкувати поведінку від суперкласу
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

У попередньому завданні ви створили `supertype`, який називався `Animal` і визначав поведінку, спільну для всіх тварин:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

Це та наступне завдання розглянуть, як повторно використовувати методи `Animal` всередині `Bird` і `Dog` без повторного їх визначення. Для цього застосовується техніка, яка називається успадкуванням. Це завдання охоплює перший крок: створити екземпляр `supertype` (або батьківського елемента). Ви вже знаєте один спосіб створити екземпляр `Animal` за допомогою оператора `new`:

```js
let animal = new Animal();
```

Існують деякі недоліки використання цього синтаксису для успадкування, які занадто складні для розгляду в цьому завданні. Замість цього ось альтернативний підхід без цих недоліків:

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)` створює новий об’єкт і встановлює `obj` як `prototype` нового об’єкта. Нагадаємо, що `prototype` — це як «рецепт» для створення об’єкта. Встановлюючи `prototype` `animal` рівним `prototype` `Animal`, ви фактично даєте екземпляру `animal` той самий «рецепт», що й будь-якому іншому екземпляру `Animal`.

```js
animal.eat();
animal instanceof Animal;
```

Метод `instanceof` тут поверне `true`.

# --instructions--

Використайте `Object.create`, щоб створити два екземпляри `Animal` з іменами `duck` і `beagle`.

# --hints--

Змінна `duck` має бути визначена.

```js
assert(typeof duck !== 'undefined');
```

Змінна `beagle` має бути визначена.

```js
assert(typeof beagle !== 'undefined');
```

Змінна `duck` має бути ініціалізована `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

Змінна `beagle` має бути ініціалізована `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

`duck` має мати `prototype` рівним `Animal`.

```js
assert(duck instanceof Animal);
```

`beagle` має мати `prototype` рівним `Animal`.

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
