---
id: 587d7daf367417b2b2512b7f
title: Змінити прототип на новий об’єкт
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

До цього моменту ви додавали властивості до `prototype` поодинці:

```js
Bird.prototype.numLegs = 2;
```

Це стає втомливим після додавання кількох властивостей.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

Більш ефективний спосіб — встановити `prototype` на новий об’єкт, який уже містить властивості. Таким чином, властивості додаються всі одразу:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

Додайте властивість `numLegs` та два методи `eat()` і `describe()` до `prototype` об’єкта `Dog`, встановивши `prototype` на новий об’єкт.

# --hints--

`Dog.prototype` має бути встановлено на новий об’єкт.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

`Dog.prototype` має містити властивість `numLegs`.

```js
assert(Dog.prototype.numLegs !== undefined);
```

`Dog.prototype` має містити метод `eat()`.

```js
assert(typeof Dog.prototype.eat === 'function');
```

`Dog.prototype` має містити метод `describe()`.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
