---
id: 587d7dae367417b2b2512b7c
title: Використання властивостей прототипу для зменшення дублювання коду
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

Оскільки `numLegs` ймовірно матиме однакове значення для всіх екземплярів `Bird`, у вас фактично є дубльована змінна `numLegs` всередині кожного екземпляра `Bird`.

Це може не бути проблемою, коли є лише два екземпляри, але уявіть, якщо їх мільйони. Це буде дуже багато дубльованих змінних.

Кращий спосіб — використовувати `prototype` `Bird`. Властивості в `prototype` спільні для ВСІХ екземплярів `Bird`. Ось як додати `numLegs` до `Bird prototype`:

```js
Bird.prototype.numLegs = 2;
```

Тепер усі екземпляри `Bird` мають властивість `numLegs`.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

Оскільки всі екземпляри автоматично мають властивості на `prototype`, уявляйте `prototype` як «рецепт» для створення об’єктів. Зверніть увагу, що `prototype` для `duck` і `canary` є частиною конструктора `Bird` як `Bird.prototype`.

# --instructions--

Додайте властивість `numLegs` до `prototype` `Dog`

# --hints--

`beagle` має мати властивість `numLegs`.

```js
assert(beagle.numLegs !== undefined);
```

`beagle.numLegs` має бути числом.

```js
assert(typeof beagle.numLegs === 'number');
```

`numLegs` має бути властивістю `prototype`, а не власною властивістю.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
