---
id: 587d7db1367417b2b2512b88
title: Заміщення успадкованих методів
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

У попередніх уроках ви дізналися, що об’єкт може успадковувати свою поведінку (методи) від іншого об’єкта, посилаючись на його `prototype` об’єкт:

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

Потім `ChildObject` отримав власні методи, додаючи їх до свого `prototype`:

```js
ChildObject.prototype.methodName = function() {...};
```

Можна заміщати успадкований метод. Це робиться так само — додаванням методу до `ChildObject.prototype` з тим самим ім’ям методу, який потрібно замістити. Ось приклад, як `Bird` заміщує метод `eat()`, успадкований від `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

Якщо у вас є екземпляр `let duck = new Bird();` і ви викликаєте `duck.eat()`, ось як JavaScript шукає метод у ланцюжку `prototype` об’єкта `duck`:

1.  `duck` => Чи визначено `eat()` тут? Ні.
2.  `Bird` => Чи визначено `eat()` тут? => Так. Виконати його і припинити пошук.
3.  `Animal` => `eat()` також визначено, але JavaScript припинив пошук раніше, ніж дійшов до цього рівня.
4.  Object => JavaScript припинив пошук раніше, ніж дійшов до цього рівня.

# --instructions--

Замістьіть метод `fly()` для `Penguin`, щоб він повертав рядок `Alas, this is a flightless bird.`

# --hints--

`penguin.fly()` має повертати рядок `Alas, this is a flightless bird.`

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

Метод `bird.fly()` має повертати рядок `I am flying!`

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
