---
id: 587d7db0367417b2b2512b83
title: "Використовуйте успадкування, щоб не повторювати код"
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

Існує принцип у програмуванні, який називається <dfn>Не повторюйся (DRY)</dfn>. Проблема повторюваного коду полягає в тому, що будь-яка зміна вимагає виправлення коду в кількох місцях. Це зазвичай означає більше роботи для програмістів і більше можливостей для помилок.

Зверніть увагу на прикладі нижче, що метод `describe` спільний для `Bird` і `Dog`:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Метод `describe` повторюється в двох місцях. Код можна відредагувати відповідно до принципу DRY, створивши `supertype` (або батьківський елемент) під назвою `Animal`:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

Оскільки `Animal` містить метод `describe`, ви можете видалити його з `Bird` і `Dog`:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

Метод `eat` повторюється і в `Cat`, і в `Bear`. Відредагуйте код у дусі DRY, перемістивши метод `eat` до `Animal` `supertype`.

# --hints--

`Animal.prototype` має мати властивість `eat`.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

`Bear.prototype` не має мати властивість `eat`.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

`Cat.prototype` не має мати властивість `eat`.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
