---
id: 587d7daf367417b2b2512b7d
title: Ітерувати над усіма пропами
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

Ви вже бачили два типи пропів: <dfn>власні пропи</dfn> та `prototype` пропи. Власні пропи визначаються безпосередньо на самому екземплярі об’єкта. А пропи прототипу визначаються на `prototype`.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

Ось як додати власні пропи `duck` до масиву `ownProps` і `prototype` пропи до масиву `prototypeProps`:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

`console.log(ownProps)` виведе `["name"]` у консоль, а `console.log(prototypeProps)` виведе `["numLegs"]`.

# --instructions--

Додайте всі власні пропи `beagle` до масиву `ownProps`. Додайте всі `prototype` пропи `Dog` до масиву `prototypeProps`.

# --hints--

Масив `ownProps` повинен містити лише `name`.

```js
assert.deepEqual(ownProps, ['name']);
```

Масив `prototypeProps` повинен містити лише `numLegs` пропи.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

Вирішіть це завдання без використання вбудованого методу `Object.keys()`.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
