---
id: 587d7dad367417b2b2512b76
title: зробіть код більш багаторазовим за допомогою ключового слова this
challengeType: 1
forumTopicId: 301321
dashedName: make-code-more-reusable-with-the-this-keyword
---

# --description--

У останньому завданні було додано метод до об’єкта `duck`. Він використовував точкову нотацію `duck.name` для доступу до значення властивості `name` всередині оператора return:

```js
sayName: function() {return "The name of this duck is " + duck.name + ".";}
```

Хоча це дійсний спосіб доступу до властивості об’єкта, тут є підводний камінь. Якщо ім’я змінної зміниться, будь-який код, що посилається на початкове ім’я, також потрібно буде оновити. У короткому визначенні об’єкта це не проблема, але якщо об’єкт має багато посилань на свої властивості, ймовірність помилки зростає.

Спосіб уникнути цих проблем — використати ключове слово `this`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + this.name + ".";}
};
```

`this` — це глибока тема, і наведений вище приклад — лише один із способів його використання. У поточному контексті `this` посилається на об’єкт, з яким пов’язаний метод: `duck`. Якщо ім’я об’єкта зміниться на `mallard`, не потрібно шукати всі посилання на `duck` у коді. Це робить код багаторазовим і легшим для читання.

# --instructions--

Змініть метод `dog.sayLegs`, щоб видалити всі посилання на `dog`. Використайте приклад `duck` як орієнтир.

# --hints--

`dog.sayLegs()` має повертати заданий рядок.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

Ваш код має використовувати ключове слово `this` для доступу до властивості `numLegs` об’єкта `dog`.

```js
assert(__helpers.removeJSComments(code).match(/this\.numLegs/g));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs: function() {return "This dog has " + dog.numLegs + " legs.";}
};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
