---
id: 5a24c314108439a4d4036151
title: "Usa un'istruzione switch per gestire azioni multiple"
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

Puoi dire all'archivio Redux come gestire tipi di azioni multiple. Supponiamo che tu stia gestendo l'autenticazione degli utenti nel tuo archivio Redux. Vuoi avere una rappresentazione dello stato per quando gli utenti sono connessi e quando sono disconnessi. Rappresenti questo con un singolo oggetto stato con la proprietà `authenticated`. Hai anche bisogno di creatori di azioni che creino azioni corrispondenti al login e al logout dell'utente, insieme agli oggetti azione stessi.

# --instructions--

L'editor del codice ha un archivio, azioni e creatori di azioni già configurati per te. Completa la funzione `reducer` per gestire azioni di autenticazione multiple. Usa un'istruzione JavaScript `switch` nel `reducer` per rispondere a diversi eventi di azione. Questo è un pattern standard nella scrittura di reducer Redux. L'istruzione switch dovrebbe fare lo switch su `action.type` e restituire lo stato di autenticazione appropriato.

**Nota:** A questo punto, non preoccuparti dell'immutabilità dello stato, dato che è piccolo e semplice in questo esempio. Per ogni azione, puoi restituire un nuovo oggetto — per esempio, `{authenticated: true}`. Inoltre, non dimenticare di scrivere un caso `default` nella tua istruzione switch che restituisca l'attuale `state`. Questo è importante perché una volta che la tua app ha più reducer, tutti vengono eseguiti ogni volta che viene effettuato un dispatch di un'azione, anche quando l'azione non è correlata a quel reducer. In tal caso, vuoi assicurarti di restituire l'attuale `state`.

# --hints--

Chiamare la funzione `loginUser` dovrebbe restituire un oggetto con la proprietà type impostata sulla stringa `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

Chiamare la funzione `logoutUser` dovrebbe restituire un oggetto con la proprietà type impostata sulla stringa `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

L'archivio dovrebbe essere inizializzato con un oggetto con una proprietà `authenticated` impostata su `false`.

```js
assert(store.getState().authenticated === false);
```

Effettuare il dispatch di `loginUser` dovrebbe aggiornare la proprietà `authenticated` nello stato dell'archivio a `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

Effettuare il dispatch di `logoutUser` dovrebbe aggiornare la proprietà `authenticated` nello stato dell'archivio a `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

La funzione `authReducer` dovrebbe gestire tipi di azioni multiple con un'istruzione `switch`.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
