---
id: 5a24c314108439a4d4036158
title: non modificare mai lo stato
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

Queste ultime sfide descrivono diversi metodi per far rispettare il principio chiave dell'immutabilità dello stato in Redux. Stato immutabile significa che non modifichi mai direttamente lo stato, ma restituisci invece una nuova copia dello stato.

Se facessi uno snapshot dello stato di un'app Redux nel tempo, vedresti qualcosa come `state 1`, `state 2`, `state 3`,`state 4`, `...` e così via, dove ogni stato può essere simile al precedente, ma ognuno è un pezzo distinto di dati. Questa immutabilità, infatti, è ciò che fornisce funzionalità come il debugging con viaggio nel tempo di cui potresti aver sentito parlare.

Redux non impone attivamente l'immutabilità dello stato nel suo archivio o nei reducer, quella responsabilità spetta al programmatore. Fortunatamente, JavaScript (specialmente ES6) offre diversi strumenti utili che puoi usare per far rispettare l'immutabilità del tuo stato, che sia un `string`, `number`, `array` o `object`. Nota che stringhe e numeri sono valori primitivi e sono immutabili per natura. In altre parole, 3 è sempre 3. Non puoi cambiare il valore del numero 3. Un `array` o `object`, invece, è mutabile. In pratica, il tuo stato probabilmente consisterà in un `array` o `object`, poiché queste sono strutture dati utili per rappresentare molti tipi di informazioni.

# --instructions--

Nel codice c'è un `store` e un `reducer` per gestire gli elementi da fare. Completa la scrittura del caso `ADD_TO_DO` nel reducer per aggiungere un nuovo elemento da fare allo stato. Ci sono diversi modi per farlo con JavaScript standard o ES6. Prova a trovare un modo per restituire un nuovo array con l'elemento da `action.todo` aggiunto in coda.

# --hints--

L'archivio Redux dovrebbe esistere e inizializzarsi con uno stato uguale all'array `todos` nel codice.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo` e `immutableReducer` dovrebbero essere entrambe funzioni.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Eseguire il dispatch di un'azione di tipo `ADD_TO_DO` sull'archivio Redux dovrebbe aggiungere un elemento `todo` e NON dovrebbe modificare lo stato.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
