---
id: 5a24c314108439a4d4036153
title: "Registrare un listener per l'archivio"
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Un altro metodo a cui puoi consultare sull'oggetto Redux `store` è `store.subscribe()`. Questo ti permette di iscrivere funzioni listener all'archivio, che vengono chiamate ogni volta che un'azione viene inviata all'archivio. Un uso semplice di questo metodo è iscrivere una funzione al tuo archivio che semplicemente registra un messaggio ogni volta che un'azione viene ricevuta e l'archivio viene aggiornato.

# --instructions--

Scrivi una funzione di callback che incrementa la variabile globale `count` ogni volta che l'archivio riceve un'azione, e passa questa funzione al metodo `store.subscribe()`. Vedrai che `store.dispatch()` viene chiamato tre volte di seguito, ogni volta passando direttamente un oggetto azione. Osserva l'output della console tra le inviate delle azioni per vedere gli aggiornamenti in atto.

# --hints--

Inviare l'azione `ADD` all'archivio dovrebbe incrementare lo stato di `1`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

Dovrebbe esserci una funzione listener iscritta all'archivio usando `store.subscribe`.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe` dovrebbe ricevere una funzione.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

La funzione passata a `store.subscribe` non dovrebbe essere chiamata.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

La callback a `store.subscribe` dovrebbe anche incrementare la variabile globale `count` mentre l'archivio viene aggiornato.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
