---
id: 5a24c314108439a4d403615b
title: Copiare un oggetto con Object.assign
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

Le ultime sfide hanno lavorato con gli array, ma ci sono modi per aiutare a garantire l'immutabilità dello stato anche quando lo stato è un `object`. Uno strumento utile per gestire gli oggetti è l'utility `Object.assign()`. `Object.assign()` prende un oggetto obiettivo e oggetti sorgente e mappa le proprietà dagli oggetti sorgente all'oggetto obiettivo. Qualsiasi proprietà corrispondente viene sovrascritta dalle proprietà negli oggetti sorgente. Questo comportamento è comunemente usato per creare copie superficiali di oggetti passando un oggetto vuoto come primo argomento seguito dall'oggetto o dagli oggetti che vuoi copiare. Ecco un esempio:

```js
const newObject = Object.assign({}, obj1, obj2);
```

Questo crea `newObject` come un nuovo `object`, che contiene le proprietà che esistono attualmente in `obj1` e `obj2`.

# --instructions--

Lo stato e le azioni di Redux sono stati modificati per gestire un `object` per il `state`. Modifica il codice per restituire un nuovo oggetto `state` per le azioni di tipo `ONLINE`, che impostano la proprietà `status` sulla stringa `online`. Prova a usare `Object.assign()` per completare la sfida.

# --hints--

L'archivio Redux dovrebbe esistere e inizializzarsi con uno stato equivalente all'oggetto `defaultState` dichiarato alla riga 1.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` e `immutableReducer` dovrebbero essere entrambe funzioni.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

L'invio di un'azione di tipo `ONLINE` dovrebbe aggiornare la proprietà `status` nello stato a `online` e NON dovrebbe mutare lo stato.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

`Object.assign` dovrebbe essere usato per restituire il nuovo stato.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
