---
id: 5a24c314108439a4d4036156
title: Usa il middleware per gestire azioni asincrone
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

Finora queste sfide hanno evitato di parlare di azioni asincrone, ma sono una parte inevitabile dello sviluppo web. A un certo punto dovrai chiamare endpoint asincroni nella tua app Redux, quindi come gestire queste richieste? Redux fornisce un middleware progettato appositamente per questo scopo, chiamato Redux Thunk middleware. Ecco una breve descrizione di come usarlo con Redux.

Per includere Redux Thunk middleware, lo passi come argomento a `Redux.applyMiddleware()`. Questa istruzione viene poi fornita come secondo parametro opzionale alla funzione `createStore()`. Dai un’occhiata al codice in fondo all’editor per vederlo. Poi, per creare un’azione asincrona, ritorni una funzione nel creatore di azioni che prende `dispatch` come argomento. All’interno di questa funzione puoi dispatchare azioni ed eseguire richieste asincrone.

In questo esempio, una richiesta asincrona è simulata con una chiamata `setTimeout()`. È comune dispatchare un’azione prima di iniziare qualsiasi comportamento asincrono così che lo stato della tua applicazione sappia che alcuni dati sono in fase di richiesta (questo stato potrebbe mostrare un’icona di caricamento, per esempio). Poi, una volta ricevuti i dati, dispatchi un’altra azione che porta i dati come payload insieme all’informazione che l’azione è completata.

Ricorda che stai passando `dispatch` come parametro a questo speciale creatore di azioni. Questo è ciò che userai per dispatchare le tue azioni, basta passare l’azione direttamente a dispatch e il middleware si occupa del resto.

# --instructions--

Scrivi entrambi i dispatch nel creatore di azioni `handleAsync()`. Dispatcha `requestingData()` prima della `setTimeout()` (la chiamata API simulata). Poi, dopo aver ricevuto i dati (finti), dispatcha l’azione `receivedData()`, passando questi dati. Ora sai come gestire azioni asincrone in Redux. Tutto il resto continua a comportarsi come prima.

# --hints--

Il creatore di azioni `requestingData` dovrebbe ritornare un oggetto di tipo uguale al valore di `REQUESTING_DATA`.

```js
assert(requestingData().type === REQUESTING_DATA);
```

Il creatore di azioni `receivedData` dovrebbe ritornare un oggetto di tipo uguale al valore di `RECEIVED_DATA`.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer` dovrebbe essere una funzione.

```js
assert(typeof asyncDataReducer === 'function');
```

Dispatchare il creatore di azioni `requestingData` dovrebbe aggiornare la proprietà fetching dello store `state` a `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

Dispatchare `handleAsync` dovrebbe dispatchare l’azione di richiesta dati e poi dispatchare l’azione di dati ricevuti dopo un ritardo.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
