---
id: 5a24c314108439a4d403614f
title: Inviare un evento di azione
challengeType: 6
forumTopicId: 301442
dashedName: dispatch-an-action-event
---

# --description--

Il metodo `dispatch` è quello che usi per inviare azioni all'archivio Redux. Chiamare `store.dispatch()` e passare il valore restituito da un creatore di azioni invia un'azione all'archivio.

Ricorda che i creatori di azioni restituiscono un oggetto con una proprietà type che specifica il tipo di azione che è avvenuta. Poi il metodo invia un oggetto azione all'archivio Redux. Basandoci sull'esempio della sfida precedente, le seguenti righe sono equivalenti e inviano entrambe l'azione di tipo `LOGIN`:

```js
store.dispatch(actionCreator());
store.dispatch({ type: 'LOGIN' });
```

# --instructions--

L'archivio Redux nel codice ha uno stato inizializzato che è un oggetto contenente una proprietà `login` attualmente impostata su `false`. C'è anche un creatore di azioni chiamato `loginAction()` che restituisce un'azione di tipo `LOGIN`. Invia l'azione `LOGIN` all'archivio Redux chiamando il metodo `dispatch` e passando l'azione creata da `loginAction()`.

# --hints--

Chiamare la funzione `loginAction` dovrebbe restituire un oggetto con la proprietà `type` impostata sulla stringa `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

L'archivio dovrebbe essere inizializzato con un oggetto con la proprietà `login` impostata su `false`.

```js
assert(store.getState().login === false);
```

Il metodo `store.dispatch()` dovrebbe essere usato per inviare un'azione di tipo `LOGIN`.

```js
assert(
    (function () {
      let noWhiteSpace = code.replace(/\s/g, '');
      return (
        noWhiteSpace.includes('store.dispatch(loginAction())') ||
        noWhiteSpace.includes("store.dispatch({type: 'LOGIN'})") === true
      );
    })()
);
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

// Dispatch the action here:
```

# --solutions--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

store.dispatch(loginAction());
```
