---
id: 5a24c314108439a4d403614c
title: "Ottenere lo stato dall'archivio Redux"
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

L'oggetto archivio Redux fornisce diversi metodi che ti permettono di interagire con esso. Per esempio, puoi recuperare lo stato corrente `state` contenuto nell'oggetto archivio Redux con il metodo `getState()`.

# --instructions--

Il codice della sfida precedente è riscritto in modo più conciso nell'editor del codice. Usa `store.getState()` per recuperare lo stato `state` dall'archivio `store` e assegnalo a una nuova variabile `currentState`.

# --hints--

L'archivio Redux dovrebbe avere un valore di 5 per lo stato iniziale.

```js
assert(store.getState() === 5);
```

Dovrebbe esistere una variabile `currentState` e dovrebbe essere assegnato lo stato corrente dell'archivio Redux.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
