---
id: 5a24c314108439a4d4036157
title: Scrivi un contatore con Redux
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

Ora hai imparato tutti i principi fondamentali di Redux! Hai visto come creare azioni e creatori di azioni, creare un archivio Redux, inviare le tue azioni all'archivio e progettare aggiornamenti di stato con reducer puri. Hai anche visto come gestire stati complessi con la composizione di reducer e gestire azioni asincrone. Questi esempi sono semplicistici, ma questi concetti sono i principi fondamentali di Redux. Se li comprendi bene, sei pronto per iniziare a costruire la tua app Redux. Le prossime sfide coprono alcuni dettagli riguardanti l'immutabilità di `state`, ma prima, ecco una revisione di tutto ciò che hai imparato finora.

# --instructions--

In questa lezione implementerai un semplice contatore con Redux da zero. Le basi sono fornite nell'editor di codice, ma dovrai completare i dettagli! Usa i nomi forniti e definisci i creatori di azioni `incAction` e `decAction`, i tipi di azione `counterReducer()`, `INCREMENT` e `DECREMENT`, e infine l'archivio Redux `store`. Una volta terminato, dovresti essere in grado di inviare azioni `INCREMENT` o `DECREMENT` per incrementare o decrementare lo stato contenuto nell'archivio `store`. Buona fortuna a costruire la tua prima app Redux!

# --hints--

Il creatore di azioni `incAction` dovrebbe restituire un oggetto azione con `type` uguale al valore di `INCREMENT`

```js
assert(incAction().type === INCREMENT);
```

Il creatore di azioni `decAction` dovrebbe restituire un oggetto azione con `type` uguale al valore di `DECREMENT`

```js
assert(decAction().type === DECREMENT);
```

Eseguire `store.getState()` dovrebbe restituire un numero

```js
assert(typeof store.getState() === 'number');
```

L'archivio Redux dovrebbe inizializzare con uno `state` pari a 0.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Inviare `incAction` all'archivio Redux dovrebbe incrementare il `state` di 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Inviare `decAction` all'archivio Redux dovrebbe decrementare il `state` di 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer` dovrebbe essere una funzione

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
