---
id: 5a24c314108439a4d4036154
title: Combinare più reducer
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

Quando lo stato della tua app inizia a diventare più complesso, può essere allettante dividere lo stato in più parti. Ricorda però il primo principio di Redux: tutto lo stato dell'app è contenuto in un unico oggetto stato nell'archivio. Perciò, Redux offre la composizione dei reducer come soluzione per un modello di stato complesso. Definisci più reducer per gestire diverse parti dello stato della tua applicazione, poi componi questi reducer insieme in un reducer radice. Il reducer radice viene poi passato al metodo Redux `createStore()`.

Per permetterci di combinare più reducer insieme, Redux fornisce il metodo `combineReducers()`. Questo metodo accetta un oggetto come argomento in cui definisci proprietà che associano chiavi a specifiche funzioni reducer. Il nome che dai alle chiavi sarà usato da Redux come nome per la parte di stato associata.

Tipicamente, è buona pratica creare un reducer per ogni parte dello stato dell'applicazione quando sono distinte o uniche in qualche modo. Per esempio, in un'app per prendere appunti con autenticazione utente, un reducer potrebbe gestire l'autenticazione mentre un altro gestisce il testo e gli appunti che l'utente invia. Per un'applicazione del genere, potremmo scrivere il metodo `combineReducers()` così:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

Ora, la chiave `notes` conterrà tutto lo stato associato ai nostri appunti e gestito dal nostro `notesReducer`. Ecco come più reducer possono essere combinati per gestire uno stato applicativo più complesso. In questo esempio, lo stato contenuto nell'archivio Redux sarebbe un unico oggetto contenente le proprietà `auth` e `notes`.

# --instructions--

Sono fornite le funzioni `counterReducer()` e `authReducer()` nell'editor del codice, insieme a un archivio Redux. Completa la scrittura della funzione `rootReducer()` usando il metodo `Redux.combineReducers()`. Assegna `counterReducer` a una chiave chiamata `count` e `authReducer` a una chiave chiamata `auth`.

# --hints--

Il `counterReducer` dovrebbe incrementare e decrementare il `state`.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

Il `authReducer` dovrebbe alternare il `state` di `authenticated` tra `true` e `false`.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

L'archivio `state` dovrebbe avere due chiavi: `count`, che contiene un numero, e `auth`, che contiene un oggetto. L'oggetto `auth` dovrebbe avere una proprietà `authenticated`, che contiene un booleano.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

Il `rootReducer` dovrebbe essere una funzione che combina il `counterReducer` e il `authReducer`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
