---
id: 5a24c314108439a4d4036159
title: "Usa l'operatore di espansione sugli array"
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Una soluzione introdotta da ES6 per aiutare a garantire l'immutabilità dello stato in Redux è l'operatore di espansione: `...`. L'operatore di espansione ha diverse applicazioni, una delle quali è particolarmente adatta alla sfida precedente di creare un nuovo array a partire da un array esistente. Si tratta di una sintassi relativamente nuova, ma molto usata. Per esempio, se hai un array `myArray` e scrivi:

```js
let newArray = [...myArray];
```

`newArray` è ora una copia di `myArray`. Entrambi gli array esistono ancora separatamente in memoria. Se esegui una mutazione come `newArray.push(5)`, `myArray` non cambia. `...` effettivamente *espande* i valori in `myArray` in un nuovo array. Per clonare un array ma aggiungere valori aggiuntivi nel nuovo array, potresti scrivere `[...myArray, 'new value']`. Questo restituirebbe un nuovo array composto dai valori in `myArray` e dalla stringa `new value` come ultimo valore. La sintassi di espansione può essere usata più volte nella composizione di array come questa, ma è importante notare che crea solo una copia superficiale dell'array. Cioè, fornisce operazioni immutabili solo per array monodimensionali.

# --instructions--

Usa l'operatore di espansione per restituire una nuova copia dello stato quando viene aggiunta una to-do.

# --hints--

L'archivio Redux dovrebbe esistere e inizializzarsi con uno stato uguale a `["Do not mutate state!"]`.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo` e `immutableReducer` dovrebbero essere entrambe funzioni.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Eseguire il dispatch di un'azione di tipo `ADD_TO_DO` sull'archivio Redux dovrebbe aggiungere un elemento `todo` e NON dovrebbe mutare lo stato.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

L'operatore di espansione dovrebbe essere usato per restituire un nuovo stato.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
