---
id: 5a24c314108439a4d403614e
title: Definire un creatore di action
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

Dopo aver creato un'action, il passaggio successivo è inviare l'action all'archivio Redux così da poter aggiornare il suo stato. In Redux, definisci i creatori di action per fare questo. Un creatore di action è semplicemente una funzione JavaScript che restituisce un'action. In altre parole, i creatori di action creano oggetti che rappresentano eventi di action.

# --instructions--

Definisci una funzione chiamata `actionCreator()` che restituisce l'oggetto `action` quando viene chiamata.

# --hints--

La funzione `actionCreator` dovrebbe esistere.

```js
assert(typeof actionCreator === 'function');
```

Chiamare la funzione `actionCreator` dovrebbe restituire l'oggetto `action`.

```js
assert(typeof action === 'object');
```

L'oggetto `action` restituito dovrebbe avere una proprietà chiave `type` con valore `LOGIN`.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
