---
id: 5a24c314108439a4d4036155
title: "Inviare dati dell'azione all'archivio"
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

Ormai hai imparato come inviare azioni all'archivio Redux, ma finora queste azioni non contenevano alcuna informazione oltre a un `type`. Puoi anche inviare dati specifici insieme alle tue azioni. Infatti, questo è molto comune perché le azioni di solito derivano da qualche interazione dell'utente e tendono a trasportare con sé dei dati. L'archivio Redux spesso ha bisogno di conoscere questi dati.

# --instructions--

Nel codice è definito un creatore di azioni `notesReducer()` di base e uno `addNoteText()`. Completa il corpo della funzione `addNoteText()` in modo che restituisca un oggetto `action`. L'oggetto dovrebbe includere una proprietà `type` con valore `ADD_NOTE` e anche una proprietà `text` impostata sui dati `note` passati al creatore di azioni. Quando chiami il creatore di azioni, passerai informazioni specifiche sulla nota a cui puoi accedere per l'oggetto.

Poi, completa la scrittura dell'istruzione `switch` nel `notesReducer()`. Devi aggiungere un caso che gestisca le azioni `addNoteText()`. Questo caso dovrebbe attivarsi ogni volta che c'è un'azione di tipo `ADD_NOTE` e dovrebbe restituire la proprietà `text` sull'oggetto `action` in arrivo come nuovo `state`.

L'azione viene inviata in fondo al codice. Quando hai finito, esegui il codice e guarda la console. È tutto ciò che serve per inviare dati specifici dell'azione all'archivio e usarli quando aggiorni lo `state` dell'archivio.

# --hints--

Il creatore di azioni `addNoteText` dovrebbe restituire un oggetto con chiavi `type` e `text`.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

Inviare un'azione di tipo `ADD_NOTE` con il creatore di azioni `addNoteText` dovrebbe aggiornare lo `state` alla stringa passata al creatore di azioni.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
