---
id: 60a3e3396c7b40068ad69975
title: Passo 11
challengeType: 0
dashedName: step-11
---

# --description--

Escreva uma nova regra usando o seletor de classe `.frame`.

Use a declaração abreviada de `border` para dar ao elemento `.frame` uma borda sólida e preta com uma largura de `50px`.

# --hints--

Você deve ter um seletor `.frame`.

```js
const hasFrame = new __helpers.CSSHelp(document).getStyle('.frame');
assert(hasFrame);
```

Você deve definir a propriedade `border` como `50px solid black`.

```js
const hasBorder = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.border === '50px solid black');
assert(hasBorder);
```

Seu elemento `.frame` deve ter uma propriedade `border` com o valor `50px solid black`.

```js
const frameBorder = new __helpers.CSSHelp(document).getStyle('.frame')?.getPropertyValue('border');
assert(frameBorder === '50px solid black');
```

# --seed--

## --seed-contents--

```css
.canvas {
  width: 500px;
  height: 600px;
  background-color: #4d0f00;
}
--fcc-editable-region--

--fcc-editable-region--
```

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
    <link href="./styles.css" rel="stylesheet">
  </head>
  <body>
    <div class="frame">
      <div class="canvas">
      </div>
    </div>
  </body>
</html>
```
