---
id: 60a3e3396c7b40068ad6996b
title: Passo 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Comece seu projeto adicionando um elemento `img` com um `src` de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png` dentro do elemento `body`.

# --hints--

O código deve ter um elemento `img`.

```js
assert(document.querySelectorAll('img').length === 1);
```

O elemento `img` deve ter para o atributo `src` o valor de `https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png`.

```js
assert(document.querySelector('img').getAttribute('src') === 'https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-1.png');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
