---
id: 60a3e3396c7b40068ad6996e
title: Passo 4
challengeType: 0
dashedName: step-4
---

# --description--

A margem é a área fora da caixa e pode ser usada para controlar o espaço entre outras caixas ou elementos.

Aqui, o elemento inferior tem uma margem superior maior, empurrando-o mais para baixo na página.

Agora que você revisou rapidamente o modelo de caixa CSS, é hora de começar a pintura Rothko.

Remova o elemento `img`.

# --hints--

Você não deve ter um elemento `img` no código.

```js
assert(document.querySelector('img') === null);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Rothko Painting</title>
  </head>
  <body>
--fcc-editable-region--
    <img src="https://cdn.freecodecamp.org/curriculum/css-box-model/diagram-3.png">
--fcc-editable-region--
  </body>
</html>
```
