---
id: 587d7fac367417b2b2512bdd
title: استخدام المقاييس الديناميكية
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

تُفيد طريقتا D3 `min()` و `max()` في ضبط المقياس.

عند وجود مجموعة بيانات معقدة، أولوية واحدة هي ضبط المقياس بحيث تتناسب الصورة البيانية مع عرض وارتفاع حاوية SVG. تريد أن تُرسم كل البيانات داخل SVG لتكون مرئية في صفحة الويب.

المثال أدناه يضبط مقياس المحور السيني لبيانات مخطط التشتت. تمرر طريقة `domain()` معلومات إلى المقياس عن قيم البيانات الخام للرسم. تعطي طريقة `range()` معلومات عن المساحة الفعلية في صفحة الويب للتصوير البياني.

في المثال، يمتد المجال من 0 إلى القيمة العظمى في المجموعة. يستخدم طريقة `max()` مع دالة استدعاء راجع تعتمد على قيم x في المصفوفات. يستخدم النطاق عرض SVG (`w`)، لكنه يتضمن أيضًا حشوًا. هذا يترك مساحة بين نقاط مخطط التشتت وحافة SVG.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

قد يكون الحشو مربكًا في البداية. تخيل المحور السيني كخط أفقي من 0 إلى 500 (قيمة العرض لـ SVG). تضمين الحشو في طريقة `range()` يجبر الرسم على البدء عند 30 على ذلك الخط (بدلاً من 0)، وينتهي عند 470 (بدلاً من 500).

# --instructions--

استخدم المتغير `yScale` لإنشاء مقياس خطي للمحور الصادي. يجب أن يبدأ المجال من الصفر ويصل إلى القيمة العظمى لـ `y` في المجموعة. يجب أن يستخدم النطاق ارتفاع SVG (`h`) ويشمل الحشو.

**ملاحظة:** تذكر أن تحافظ على الرسم في الوضع الصحيح. عند ضبط النطاق لإحداثيات y، تكون القيمة الأعلى (الارتفاع ناقص الحشو) هي المعلمة الأولى، والقيمة الأقل هي المعلمة الثانية.

# --hints--

النص في `h2` يجب أن يكون `30`.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

يجب أن يكون `domain()` لـ yScale مكافئًا لـ `[0, 411]`.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

يجب أن يكون `range()` لـ yScale مكافئًا لـ `[470, 30]`.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
