---
id: 587d7fac367417b2b2512bdb
title: تعيين مجال ونطاق على مقياس
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

افتراضيًا، تستخدم المقاييس العلاقة الهوية. هذا يعني أن قيمة المُدخَل تُطابق قيمة المخرَج. مع ذلك، يمكن للمقاييس أن تكون أكثر مرونة وإثارة للاهتمام.

لنفترض أن مجموعة بيانات تحتوي على قيم تتراوح بين 50 و480. هذه هي معلومات المُدخَل للمقياس، والمعروفة أيضًا باسم <dfn>المجال</dfn>.

تريد تعيين تلك النقاط على طول المحور `x` في SVG، بين 10 وحدات و500 وحدة. هذه هي معلومات المخرَج، والمعروفة أيضًا باسم <dfn>النطاق</dfn>.

تُعيّن الطريقتان `domain()` و`range()` هذه القيم للمقياس. كلتا الطريقتين تأخذان مصفوفة تحتوي على عنصرين على الأقل كمعلمة. إليك مثالًا:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

بالترتيب، ستُعرض القيم التالية في وحدة التحكم: `10`، `500`، `323.37`، و`807.67`.

لاحظ أن المقياس يستخدم العلاقة الخطية بين قيم المجال والنطاق ليحدد ما يجب أن تكون عليه قيمة المخرَج لرقم معين. القيمة الحد الأدنى في المجال (50) تُطابق القيمة الحد الأدنى (10) في النطاق.

# --instructions--

أنشئ مقياسًا وحدد مجاله إلى `[250, 500]` ونطاقه إلى `[10, 150]`.

**ملاحظة:** يمكنك ربط الطريقتين `domain()` و`range()` على متغير `scale`.

# --hints--

يجب أن يستخدم شفرتك الطريقة `domain()`.

```js
assert.match(code, /\.domain/g);
```

يجب تعيين `domain()` للـ `scale` إلى `[250, 500]`.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

يجب أن يستخدم شفرتك الطريقة `range()`.

```js
assert.match(code, /\.range/g);
```

يجب تعيين `range()` للـ `scale` إلى `[10, 150]`.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

يجب أن يكون النص في `h2` هو `-102`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
