---
id: 587d7fa7367417b2b2512bc7
title: تغيير الأنماط بناءً على البيانات
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3 يتعلق بتصوير البيانات وعرضها. من المحتمل أن ترغب في تغيير تنسيق العناصر بناءً على البيانات.
على سبيل المثال، قد ترغب في تلوين نقطة بيانات باللون الأزرق إذا كانت قيمتها أقل من 20، وباللون الأحمر خلاف ذلك. يمكنك استخدام دالة استدعاء راجع في طريقة `style()` وتضمين المنطق الشرطي. تستخدم دالة الاستدعاء الراجع المعلمة `d` لتمثيل نقطة البيانات:

```js
selection.style('color', d => {});
```

طريقة `style()` ليست مقصورة على تعيين `color` فقط - يمكن استخدامها مع خصائص CSS أخرى أيضًا.

# --instructions--

أضف طريقة `style()` إلى الشفرة في المحرر لتعيين `color` لعناصر `h2` بشكل شرطي. اكتب دالة الاستدعاء الراجع بحيث إذا كانت قيمة البيانات أقل من 20، تُرجع اللون الأحمر، وإلا تُرجع اللون الأخضر.

**ملاحظة:** يمكنك استخدام منطق if-else، أو المعامل الثلاثي.

# --hints--

يجب أن يكون للعنصر `h2` الأول `color` باللون الأحمر.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب أن يكون للعنصر `h2` الثاني `color` باللون الأخضر.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

يجب أن يكون للعنصر `h2` الثالث `color` باللون الأخضر.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

يجب أن يكون للعنصر `h2` الرابع `color` باللون الأحمر.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب أن يكون للعنصر `h2` الخامس `color` باللون الأخضر.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

يجب أن يكون للعنصر `h2` السادس `color` باللون الأحمر.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب أن يكون للعنصر `h2` السابع `color` باللون الأخضر.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

يجب أن يكون للعنصر `h2` الثامن `color` باللون الأحمر.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

يجب أن يكون للعنصر `h2` التاسع `color` باللون الأحمر.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
