---
id: 587d7faa367417b2b2512bd2
title: إضافة مسميات لعناصر D3
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

تتيح لك D3 تسمية عنصر في الرسم البياني، مثل شريط، باستخدام عنصر SVG `text`.

مثل عنصر `rect`، يحتاج عنصر `text` إلى خاصيتي `x` و `y` لوضعه على SVG. كما يحتاج إلى وصول للبيانات لعرض تلك القيم.

تمنحك D3 تحكمًا عالي المستوى في كيفية تسمية أشرطة الرسم البياني.

# --instructions--

الشيفرة في المحرر تربط البيانات بالفعل بكل عنصر `text` جديد. أولًا، أضف عقد `text` إلى `svg`. بعد ذلك، أضف خاصيتي الإحداثيات `x` و `y`. يجب حسابهما بنفس طريقة عناصر `rect`، لكن قيمة `y` لعنصر `text` يجب أن تجعل المسمى أعلى بمقدار 3 وحدات من الشريط. أخيرًا، استخدم طريقة D3 `text()` لتعيين المسمى مساويًا لقيمة نقطة البيانات.

**ملاحظة:** لجعل المسمى أعلى من الشريط، قرر إذا كانت قيمة `y` لعنصر `text` يجب أن تكون أكبر أو أقل بمقدار 3 من قيمة `y` للشريط.

# --hints--

يجب أن يكون للمسمى في أول عنصر `text` القيمة `12` وقيمة `y` تساوي `61`.

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

يجب أن يكون للمسمى في ثاني عنصر `text` القيمة `31` وقيمة `y` تساوي `4`.

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

يجب أن يكون للمسمى في ثالث عنصر `text` القيمة `22` وقيمة `y` تساوي `31`.

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

يجب أن يكون للمسمى في رابع عنصر `text` القيمة `17` وقيمة `y` تساوي `46`.

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

يجب أن يكون للمسمى في خامس عنصر `text` القيمة `25` وقيمة `y` تساوي `22`.

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

يجب أن يكون للمسمى في سادس عنصر `text` القيمة `18` وقيمة `y` تساوي `43`.

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

يجب أن يكون للمسمى في سابع عنصر `text` القيمة `29` وقيمة `y` تساوي `10`.

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

يجب أن يكون للمسمى في ثامن عنصر `text` القيمة `14` وقيمة `y` تساوي `55`.

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

يجب أن يكون للمسمى في تاسع عنصر `text` القيمة `9` وقيمة `y` تساوي `70`.

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```
