---
id: 587d7faa367417b2b2512bd4
title: إضافة تأثير المرور بالمؤشر إلى عنصر D3
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

يمكنك إضافة تأثيرات تبرز الشريط عندما يمر المستخدم بالمؤشر فوقه. حتى الآن، يُطبق التنسيق على المستطيلات باستخدام الطرق المدمجة في D3 وSVG، لكن يمكنك استخدام CSS أيضًا.

تحدد صنف CSS على عناصر SVG باستخدام الطريقة `attr()`. ثم تحتفظ الفئة المزيفة `:hover` للصنف الجديد بقواعد التنسيق لأي تأثيرات مرور بالمؤشر.

# --instructions--

استخدم الطريقة `attr()` لإضافة صنف `bar` إلى جميع عناصر `rect`. هذا يغير لون التعبئة `fill` للشريط إلى البني عندما تمرر المؤشر فوقه.

# --hints--

يجب أن تحتوي عناصر `rect` على صنف `bar`.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```
