---
id: 587d7fad367417b2b2512bdf
title: إضافة محاور إلى التصور
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

طريقة أخرى لتحسين مخطط التشتت هي إضافة محور x ومحور y.

يوفر D3 طريقتين، `axisLeft()` و `axisBottom()`، لعرض محور y ومحور x على التوالي. إليك مثالًا لإنشاء محور x بناءً على `xScale` في التحديات السابقة:

```js
const xAxis = d3.axisBottom(xScale);
```

الخطوة التالية هي عرض المحور على SVG. للقيام بذلك، يمكنك استخدام مكوِّن SVG عام، وهو عنصر `g`. الحرف `g` يرمز إلى مجموعة. على عكس `rect` و `circle` و `text`، المحور هو مجرد خط مستقيم عند عرضه. وبما أنه شكل بسيط، فإن استخدام `g` مناسب. الخطوة الأخيرة هي تطبيق خاصية `transform` لوضع المحور في المكان الصحيح على SVG. وإلا، سيُعرض الخط على حدود SVG ولن يكون مرئيًا. يدعم SVG أنواعًا مختلفة من `transforms`، لكن وضع المحور يحتاج إلى `translate`. عند تطبيقه على عنصر `g`، يحرك المجموعة كلها يمينًا وأسفل بمقدار القيم المعطاة. إليك مثالًا:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

الكود أعلاه يضع محور x في أسفل SVG. ثم يُمرر كمعلمة إلى طريقة `call()`. يعمل محور y بنفس الطريقة، باستثناء أن معلمة `translate` تكون على شكل `(x, 0)`. وبما أن `translate` سلسلة نصية في طريقة `attr()` أعلاه، يمكنك استخدام الربط لتضمين قيم متغيرة كمعلمات.

# --instructions--

مخطط التشتت يحتوي الآن على محور x. أنشئ محور y في متغير باسم `yAxis` باستخدام طريقة `axisLeft()`. ثم اعرض المحور باستخدام عنصر `g`. تأكد من استخدام خاصية `transform` لترجمة المحور بمقدار وحدات الحشو يمينًا، و `0` وحدات للأسفل. تذكر أن تستدعي المحور باستخدام `call()`.

# --hints--

يجب أن يستخدم شفرتك طريقة `axisLeft()` مع تمرير `yScale` كمعلمة.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

يجب أن يحتوي عنصر `g` لمحور y على خاصية `transform` لترجمة المحور بمقدار `(60, 0)`.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

يجب أن تستدعي شفرتك `yAxis`.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
