---
id: 587d7fa9367417b2b2512bcf
title: تغيير ارتفاع كل شريط ديناميكيًا
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

يمكن تعيين ارتفاع كل شريط إلى قيمة نقطة البيانات في المصفوفة، تمامًا كما تم تعيين قيمة `x` ديناميكيًا.

```js
selection.attr('property', (d, i) => {});
```

هنا `d` ستكون قيمة نقطة البيانات، و`i` سيكون فهرس نقطة البيانات في المصفوفة.

# --instructions--

غيّر دالة الاستدعاء الراجع للخاصية `height` لتُرجع قيمة البيانات مضروبة في 3.

**ملاحظة:** تذكّر أن ضرب جميع نقاط البيانات بنفس الثابت يُكبير البيانات (مثل التكبير). هذا يساعد على رؤية الفروقات بين قيم الأعمدة في هذا المثال.

# --hints--

يجب أن يكون للعنصر الأول `rect` قيمة `height` تساوي `36`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

يجب أن يكون للعنصر الثاني `rect` قيمة `height` تساوي `93`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

يجب أن يكون للعنصر الثالث `rect` ارتفاع `height` يساوي `66`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

يجب أن يكون للعنصر الرابع `rect` ارتفاع `height` يساوي `51`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

يجب أن يكون للعنصر `rect` الخامس ارتفاع `height` يساوي `75`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

يجب أن يكون للعنصر السادس `rect` ارتفاع `height` مقداره `54`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

يجب أن يكون للعنصر `rect` السابع قيمة `height` تساوي `87`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

يجب أن يكون للعنصر الثامن `rect` قيمة `height` تساوي `42`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

يجب أن يكون للعنصر `rect` التاسع ارتفاع `height` بقيمة `27`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
