---
id: 587d7fa8367417b2b2512bca
title: تغيير عرض الثميل البياني بالأعمدة
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

التحدٍّ الأخير أنشأ ثميلًا بيانيًا بالأعمدة، لكن هناك بعض تغييرات التنسيق التي يمكن أن تحسنه:

1. أضف مسافة بين كل عمود لفصلها بصريًا، ويتم ذلك بإضافة هامش إلى CSS الخاص بـ `bar` صنف

2. زد ارتفاع الأعمدة لتوضيح الفرق في القيم بشكل أفضل، ويتم ذلك بضرب القيمة في رقم لتغيير مقياس الارتفاع

# --instructions--

أولًا، أضف `margin` بقيمة `2px` إلى صنف `bar` داخل وسم `style`. بعد ذلك، غيّر دالة الاستدعاء الراجع في طريقة `style()` بحيث تُرجع قيمة تساوي `10` أضعاف القيمة الأصلية للبيانات (مع إضافة `px`).

**ملاحظة:** ضرب كل نقطة بيانات بنفس الثابت فقط يغير المقياس. يشبه التكبير، ولا يغير معنى البيانات الأساسية.

# --hints--

يجب أن يكون للعنصر `div` الأول `height` بقيمة `120` بكسل و`margin` بقيمة `2` بكسل.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` الثاني `height` بقيمة `310` بكسل و`margin` بقيمة `2` بكسل.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` الثالث `height` بقيمة `220` بكسل و`margin` بقيمة `2` بكسل.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` الرابع `height` بقيمة `170` بكسل و`margin` بقيمة `2` بكسل.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` الخامس `height` بقيمة `250` بكسل و`margin` بقيمة `2` بكسل.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` السادس `height` بقيمة `180` بكسل و`margin` بقيمة `2` بكسل.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` السابع `height` بقيمة `290` بكسل و`margin` بقيمة `2` بكسل.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` الثامن `height` بقيمة `140` بكسل و`margin` بقيمة `2` بكسل.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

يجب أن يكون للعنصر `div` التاسع `height` بقيمة `90` بكسل و`margin` بقيمة `2` بكسل.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
