---
id: 587d7fab367417b2b2512bd7
title: إنشاء مخطط انتشار باستخدام دوائر SVG
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

مخطط الانتشار هو نوع آخر من التصوير البياني. عادةً ما يستخدم دوائر لرسم نقاط البيانات، ولكل نقطة قيمتان. ترتبط هذه القيم بمحوري `x` و `y`، وتُستخدم لتحديد موقع الدائرة في التصوير البياني.

يحتوي SVG على وسم `circle` لإنشاء شكل الدائرة. يعمل هذا الوسم بطريقة مشابهة لعناصر `rect` التي استخدمتها في الثميل البياني بالأعمدة.

# --instructions--

استخدم الطرق `data()`, `enter()`, و `append()` لربط `dataset` بعناصر `circle` جديدة تُضاف إلى SVG.

**ملاحظة:** الدوائر لن تكون مرئية لأننا لم نحدد خصائصها بعد. سنقوم بذلك في التحدي التالي.

# --hints--

يجب أن يحتوي شفرتك على 10 عناصر `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
