---
id: 587d7fab367417b2b2512bda
title: إنشاء مقياس خطي باستخدام D3
challengeType: 6
forumTopicId: 301483
dashedName: create-a-linear-scale-with-d3
---

# --description--

كل من مخططات الأعمدة والنقاط المبعثرة رسمت البيانات مباشرة على SVG. مع ذلك، إذا كان ارتفاع عمود أو إحدى نقاط البيانات أكبر من قيم ارتفاع أو عرض SVG، فسيتجاوز مساحة SVG.

في D3، توجد مقاييس تساعد على رسم البيانات. `scales` هي دوال تخبر البرنامج كيف يربط مجموعة من نقاط البيانات الخام ببكسلات SVG.

على سبيل المثال، لنفترض أن لديك SVG بحجم 100x500 وتريد رسم الناتج المحلي الإجمالي (GDP) لعدد من الدول. ستكون مجموعة الأرقام في نطاق المليارات أو التريليونات. تزود D3 بنوع من المقياس ليخبره كيف يضع قيم الناتج المحلي الإجمالي الكبيرة ضمن تلك المساحة 100x500.

من غير المحتمل أن ترسم البيانات الخام كما هي. قبل الرسم، تضبط المقياس لمجموعة بياناتك كاملة، بحيث تتناسب قيم `x` و `y` مع عرض وارتفاع SVG.

لدى D3 عدة أنواع من المقاييس. للمقياس الخطي (يُستخدم عادة مع البيانات الكمية)، هناك طريقة D3 `scaleLinear()`:

```js
const scale = d3.scaleLinear();
```

افتراضيًا، يستخدم المقياس علاقة الهوية. قيمة المُدخَل هي نفسها قيمة المخرَج. يغطي تحدٍّ منفصل كيفية تغيير هذا.

# --instructions--

غيّر متغير `scale` لإنشاء مقياس خطي. ثم اضبط متغير `output` ليكون ناتج استدعاء المقياس مع معلمة مُدخَل قيمتها `50`.

# --hints--

النص في عنصر `h2` يجب أن يكون `50`.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '50');
```

يجب أن يستخدم شفرتك طريقة `scaleLinear()`.

```js
assert.match(code, /\.scaleLinear/g);
```

يجب أن يستدعي متغير `output` الدالة `scale` مع معلمة قيمتها `50`.

```js
assert.strictEqual(output, 50);
assert.match(code, /scale\(\s*50\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line

    const scale = undefined; // Create the scale here
    const output = scale(); // Call scale with an argument here

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    const output = scale(50);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
