---
id: 587d7fa6367417b2b2512bc2
title: إضافة عناصر إلى المستند باستخدام D3
challengeType: 6
forumTopicId: 301474
dashedName: add-document-elements-with-d3
---

# --description--

يحتوي D3 على عدة طرق تتيح لك إضافة وتغيير العناصر في مستندك.

تختار الطريقة `select()` عنصرًا واحدًا من المستند. تأخذ معلمة لاسم العنصر الذي تريده وتُرجع عقدة HTML لأول عنصر في المستند يطابق الاسم. إليك مثالًا:

```js
const anchor = d3.select('a');
```

يجد المثال أعلاه أول وسم الرابط على الصفحة ويحفظ عقدة HTML له في المتغير `anchor`. يمكنك استخدام التحديد مع طرق أخرى. الجزء `d3` في المثال هو مرجع لكائن D3، وهو الطريقة التي تصل بها إلى طرق D3.

هناك طريقتان مفيدتان أخريان هما `append()` و `text()`.

تأخذ الطريقة `append()` معلمة للعنصر الذي تريد إضافته إلى المستند. تُلحق عقدة HTML إلى عنصر محدد، وتُرجع مقبضًا لتلك العقدة.

تقوم الطريقة `text()` إما بتعيين نص العقدة المحددة، أو بالحصول على النص الحالي. لتعيين القيمة، تمرر سلسلة نصية كمعلمة داخل أقواس الطريقة.

إليك مثالًا يختار قائمة غير مرتبة، ويُلحق عنصر قائمة، ويضيف نصًا:

```js
d3.select('ul')?.append('li').text('Very important item');
```

يسمح لك D3 بربط عدة طرق معًا باستخدام النقاط لتنفيذ عدد من الإجراءات على التوالي.

# --instructions--

استخدم طريقة `select` لاختيار وسم `body` في المستند. ثم أضف وسم `h1` إليه باستخدام `append`، وأدخل النص `Learning D3` داخل عنصر `h1`.

# --hints--

يجب أن يحتوي `body` على عنصر `h1` واحد.

```js
const body = document.querySelector('body');
const headers = body?.querySelectorAll('h1');
assert.lengthOf(headers, 1);
```

يجب أن يحتوي عنصر `h1` على النص `Learning D3`.

```js
assert.strictEqual(document.querySelector('h1')?.textContent, 'Learning D3');
```

يجب أن يصل كودك إلى كائن `d3`.

```js
assert.match(code, /d3/g);
```

يجب أن يستخدم كودك الطريقة `select`.

```js
assert.match(code, /\.select/g);
```

يجب أن يستخدم كودك الطريقة `append`.

```js
assert.match(code, /\.append/g);
```

يجب أن يستخدم كودك الطريقة `text`.

```js
assert.match(code, /\.text/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    d3.select('body').append('h1').text('Learning D3');
  </script>
</body>
```
