---
id: 587d7fa8367417b2b2512bcc
title: عرض الأشكال باستخدام SVG
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

التحدي الأخير أنشأ عنصر `svg` بعرض وارتفاع محددين، وكان مرئيًا لأنه وُضِع له `background-color` في وسم `style`. الشفرة خصصت مساحة للعرض والارتفاع المحددين.

الخطوة التالية هي إنشاء شكل لوضعه في مساحة `svg`. هناك عدة أشكال مدعومة في SVG، مثل المستطيلات والدوائر. تُستخدم هذه الأشكال لعرض البيانات. على سبيل المثال، يمكن لشكل مستطيل (`<rect>`) في SVG أن يُنشئ عمودًا في الثميل البياني بالأعمدة.

عند وضع شكل في مساحة `svg`، يمكنك تحديد موقعه باستخدام إحداثيات `x` و `y`. نقطة الأصل (0, 0) تكون في الزاوية العلوية اليسرى. القيم الموجبة لـ `x` تدفع الشكل إلى اليمين، والقيم الموجبة لـ `y` تدفع الشكل إلى الأسفل من نقطة الأصل.

لوضع شكل في منتصف مساحة `svg` التي أبعادها 500 (عرض) × 100 (ارتفاع) من التحدي السابق، ستكون إحداثية `x` هي 250 وإحداثية `y` هي 50.

لـ `rect` في SVG أربع خصائص. هناك إحداثيات `x` و `y` لتحديد مكانه في مساحة `svg`. كما له خاصيتا `height` و `width` لتحديد الحجم.

# --instructions--

أضف شكل `rect` إلى `svg` باستخدام `append()`، وأعطه خاصية `width` بقيمة `25` وخاصية `height` بقيمة `100`. أيضًا، عيّن خاصيتي `x` و `y` في `rect` إلى `0` لكل منهما.

# --hints--

يجب أن يحتوي مستندك على عنصر `rect` واحد.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

يجب أن يحتوي عنصر `rect` على خاصية `width` بقيمة `25`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

يجب أن يحتوي عنصر `rect` على خاصية `height` بقيمة `100`.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

يجب أن يحتوي عنصر `rect` على خاصية `x` بقيمة `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

يجب أن يحتوي عنصر `rect` على خاصية `y` بقيمة `0`.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
