---
id: 587d7fab367417b2b2512bd8
title: إضافة خصائص إلى عناصر الدائرة
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

التحدي الأخير أنشأ عناصر `circle` لكل نقطة في `dataset`، وألحقها بـ SVG. لكن D3 يحتاج إلى مزيد من المعلومات عن موقع وحجم كل `circle` ليعرضها بشكل صحيح.

لـ `circle` في SVG ثلاث خصائص رئيسية. خاصيتا `cx` و `cy` هما الإحداثيات. تخبران D3 أين يضع _مركز_ الشكل على SVG. نصف القطر (خاصية `r`) يحدد حجم الـ `circle`.

تمامًا مثل إحداثي `y` في `rect`، خاصية `cy` لـ `circle` تُقاس من أعلى SVG، وليس من الأسفل.

يمكن لكل الخصائص الثلاث استخدام دالة استدعاء راجع لتعيين قيمها ديناميكيًا. تذكر أن كل الطرق المرتبطة بعد `data(dataset)` تُنفذ مرة لكل عنصر في `dataset`. المعلمة `d` في دالة الاستدعاء الراجعة تشير إلى العنصر الحالي في `dataset`، وهو مصفوفة لكل نقطة. تستخدم ترميز الأقواس مثل `d[0]` للوصول إلى القيم في تلك المصفوفة.

# --instructions--

أضف خصائص `cx` و `cy` و `r` إلى عناصر `circle`. يجب أن تكون قيمة `cx` هي الرقم الأول في المصفوفة لكل عنصر في `dataset`. يجب أن تعتمد قيمة `cy` على الرقم الثاني في المصفوفة، لكن تأكد من عرض المخطط بشكل صحيح وليس مقلوبًا. يجب أن تكون قيمة `r` هي `5` لكل الدوائر.

# --hints--

يجب أن يحتوي شفرتك على 10 عناصر `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

يجب أن يحتوي أول عنصر `circle` على قيمة `cx` تساوي `34`، وقيمة `cy` تساوي `422`، وقيمة `r` تساوي `5`.

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

يجب أن يحتوي ثاني عنصر `circle` على قيمة `cx` تساوي `109`، وقيمة `cy` تساوي `220`، وقيمة `r` تساوي `5`.

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

يجب أن يحتوي ثالث عنصر `circle` على قيمة `cx` تساوي `310`، وقيمة `cy` تساوي `380`، وقيمة `r` تساوي `5`.

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

يجب أن يحتوي رابع عنصر `circle` على قيمة `cx` تساوي `79`، وقيمة `cy` تساوي `89`، وقيمة `r` تساوي `5`.

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

يجب أن يحتوي خامس عنصر `circle` على قيمة `cx` تساوي `420`، وقيمة `cy` تساوي `280`، وقيمة `r` تساوي `5`.

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

يجب أن يحتوي سادس عنصر `circle` على قيمة `cx` تساوي `233`، وقيمة `cy` تساوي `355`، وقيمة `r` تساوي `5`.

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

يجب أن يحتوي سابع عنصر `circle` على قيمة `cx` تساوي `333`، وقيمة `cy` تساوي `404`، وقيمة `r` تساوي `5`.

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

يجب أن يحتوي ثامن عنصر `circle` على قيمة `cx` تساوي `222`، وقيمة `cy` تساوي `167`، وقيمة `r` تساوي `5`.

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

يجب أن يحتوي تاسع عنصر `circle` على قيمة `cx` تساوي `78`، وقيمة `cy` تساوي `180`، وقيمة `r` تساوي `5`.

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

يجب أن يحتوي عاشر عنصر `circle` على قيمة `cx` تساوي `21`، وقيمة `cy` تساوي `377`، وقيمة `r` تساوي `5`.

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
