---
id: 587d7fa7367417b2b2512bc6
title: إضافة تنسيق مضمَّن للعناصر
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

تتيح لك مكتبة D3 إضافة أنماط CSS مضمَّنة على العناصر الديناميكية باستخدام الطريقة `style()`.

تأخذ الطريقة `style()` زوجًا من المفتاح والقيمة مفصولًا بفاصلة كمعلمة. إليك مثالًا لتعيين لون نص التحديد إلى الأزرق:

```js
selection.style('color', 'blue');
```

# --instructions--

أضف الطريقة `style()` إلى الشفرة في المحرر لجعل كل النص المعروض يستخدم `font-family` بقيمة `verdana`.

# --hints--

يجب أن تحتوي عناصر `h2` على `font-family` بقيمة `verdana`.

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

يجب أن تستخدم شفرتك الطريقة `style()`.

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
