---
id: 587d7fac367417b2b2512bde
title: استخدام مقياس معرف مسبقًا لوضع العناصر
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

بعد إعداد المقاييس، حان الوقت لرسم مخطط التشتت مرة أخرى. المقاييس تشبه دوال المعالجة التي تحول بيانات `x` و `y` الخام إلى قيم تناسب وتُعرض بشكل صحيح على SVG. تحافظ على البيانات ضمن مساحة الرسم على الشاشة.

تُعيّن قيم خاصية الإحداثيات لشكل SVG باستخدام دالة التحجيم. يشمل ذلك خاصيتي `x` و`y` لعناصر `rect` أو `text`، أو `cx` و`cy` لعناصر `circles`. إليك مثالًا:

```js
shape.attr('x', d => xScale(d[0]));
```

تحدد المقاييس خصائص إحداثيات الشكل لوضع نقاط البيانات على SVG. لا تحتاج إلى تطبيق المقاييس عند عرض القيمة الفعلية للبيانات، مثلًا في طريقة `text()` لتلميح أو مسمى.

# --instructions--

استخدم `xScale` و`yScale` لوضع كل من أشكال `circle` و`text` على SVG. بالنسبة لـ`circles`، طبّق المقاييس لتعيين خاصيتي `cx` و`cy`. امنحها أيضًا نصف قطر مقداره `5` وحدات.

بالنسبة لعناصر `text`، طبق المقاييس لتعيين خصائص `x` و `y`. يجب أن تكون المسميات مزيحة إلى يمين النقاط. للقيام بذلك، أضف `10` وحدات إلى قيمة بيانات `x` قبل تمريرها إلى `xScale`.

# --hints--

يجب أن يحتوي شفرتك على 10 عناصر `circle`.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

يجب أن يحتوي أول عنصر `circle` على قيمة `cx` تقارب `91` وقيمة `cy` تقارب `368` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي ثاني عنصر `circle` على قيمة `cx` تقارب `159` وقيمة `cy` تقارب `181` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي ثالث عنصر `circle` على قيمة `cx` تقارب `340` وقيمة `cy` تقارب `329` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي رابع عنصر `circle` على قيمة `cx` تقارب `131` وقيمة `cy` تقارب `60` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي خامس عنصر `circle` على قيمة `cx` تقارب `440` وقيمة `cy` تقارب `237` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي سادس عنصر `circle` على قيمة `cx` تقارب `271` وقيمة `cy` تقارب `306` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي سابع عنصر `circle` على قيمة `cx` تقارب `361` وقيمة `cy` تقارب `351` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي ثامن عنصر `circle` على قيمة `cx` تقارب `261` وقيمة `cy` تقارب `132` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي تاسع عنصر `circle` على قيمة `cx` تقارب `131` وقيمة `cy` تقارب `144` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي عاشر عنصر `circle` على قيمة `cx` تقارب `79` وقيمة `cy` تقارب `326` بعد تطبيق المقاييس. كما يجب أن يكون له قيمة `r` تساوي `5`.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

يجب أن يحتوي شفرتك على 10 عناصر `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

يجب أن يحتوي أول مسمى على قيمة `x` تقارب `100` وقيمة `y` تقارب `368` بعد تطبيق المقاييس.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

يجب أن يحتوي ثاني مسمى على قيمة `x` تقارب `168` وقيمة `y` تقارب `181` بعد تطبيق المقاييس.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

يجب أن يحتوي ثالث مسمى على قيمة `x` تقارب `350` وقيمة `y` تقارب `329` بعد تطبيق المقاييس.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

يجب أن يحتوي رابع مسمى على قيمة `x` تقارب `141` وقيمة `y` تقارب `60` بعد تطبيق المقاييس.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

يجب أن يحتوي خامس مسمى على قيمة `x` تقارب `449` وقيمة `y` تقارب `237` بعد تطبيق المقاييس.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

يجب أن يحتوي سادس مسمى على قيمة `x` تقارب `280` وقيمة `y` تقارب `306` بعد تطبيق المقاييس.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

يجب أن يحتوي سابع مسمى على قيمة `x` تقارب `370` وقيمة `y` تقارب `351` بعد تطبيق المقاييس.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

يجب أن يحتوي ثامن مسمى على قيمة `x` تقارب `270` وقيمة `y` تقارب `132` بعد تطبيق المقاييس.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

يجب أن يحتوي تاسع مسمى على قيمة `x` تقارب `140` وقيمة `y` تقارب `144` بعد تطبيق المقاييس.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

يجب أن يحتوي عاشر مسمى على قيمة `x` تقارب `88` وقيمة `y` تقارب `326` بعد تطبيق المقاييس.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
