---
id: 587d7fa7367417b2b2512bc8
title: إضافة أصناف باستخدام D3
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

يصبح من الصعب إدارة الكثير من الأنماط المضمَّنة على عناصر HTML، حتى في التطبيقات الصغيرة. من الأسهل إضافة صنف إلى العناصر وتنسيق هذا الصنف مرة واحدة باستخدام قواعد CSS. يحتوي D3 على طريقة `attr()` لإضافة أي خاصية HTML إلى عنصر، بما في ذلك اسم الصنف.

تعمل طريقة `attr()` بنفس طريقة `style()`. تأخذ قيمًا مفصولة بفواصل، ويمكنها استخدام دالة استدعاء راجع. إليك مثالًا لإضافة صنف `container` إلى اختيار:

```js
selection.attr('class', 'container');
```

لاحظ أن معلمة `class` ستظل كما هي كلما احتجت إلى إضافة صنف، ولن تتغير سوى معلمة `container`.

# --instructions--

أضف طريقة `attr()` إلى الشفرة في المحرر وضع صنف `bar` على عناصر `div`.

# --hints--

يجب أن تحتوي عناصر `div` على صنف `bar`.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

يجب أن تستخدم شفرتك طريقة `attr()`.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
