---
id: 587d7fa8367417b2b2512bc9
title: تحديث ارتفاع عنصر بشكل ديناميكي
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

التحديات السابقة تناولت كيفية عرض البيانات من مصفوفة وكيفية إضافة أصناف CSS. يمكنك دمج هذه الدروس لإنشاء ثميل بياني بالأعمدة بسيط. هناك خطوتان لذلك:

1. أنشئ `div` لكل نقطة بيانات في المصفوفة

2. أعطِ كل `div` ارتفاعًا ديناميكيًا، باستخدام دالة استدعاء راجع في طريقة `style()` التي تضبط الارتفاع مساويًا لقيمة البيانات

تذكر صيغة ضبط نمط باستخدام دالة استدعاء راجع:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

أضف طريقة `style()` إلى الشفرة في المحرر لضبط خاصية `height` لكل عنصر. استخدم دالة استدعاء راجع لإرجاع قيمة نقطة البيانات مع إضافة السلسلة النصية `px` إليها.

# --hints--

يجب أن يكون للعنصر `div` الأول `height` يساوي `12` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

يجب أن يكون للعنصر `div` الثاني `height` يساوي `31` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

يجب أن يكون للعنصر `div` الثالث `height` يساوي `22` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

يجب أن يكون للعنصر `div` الرابع `height` يساوي `17` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

يجب أن يكون للعنصر `div` الخامس `height` يساوي `25` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

يجب أن يكون للعنصر `div` السادس `height` بقيمة `18` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

يجب أن يكون للعنصر `div` السابع `height` بقيمة `29` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

يجب أن يكون للعنصر `div` الثامن `height` يساوي `14` بكسل.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

يجب أن يكون للعنصر `div` التاسع `height` يساوي `9` بكسلات.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
