---
id: 587d7fa6367417b2b2512bc3
title: اختيار مجموعة من العناصر باستخدام D3
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

يحتوي D3 أيضًا على الطريقة `selectAll()` لاختيار مجموعة من العناصر. تُرجع مصفوفة من عقد HTML لكل العناصر في المستند التي تطابق سلسلة المُدخَل. إليك مثالًا لاختيار كل وسم الرابط في المستند:

```js
const anchors = d3.selectAll('a');
```

مثل الطريقة `select()`، تدعم `selectAll()` التسلسل، ويمكنك استخدامها مع طرق أخرى.

# --instructions--

اختر كل الوسوم `li` في المستند، وغيّر نصها إلى السلسلة النصية `list item` عبر تسلسل الطريقة `.text()`.

# --hints--

يجب أن يكون هناك 3 عناصر `li` في الصفحة، ويجب أن يكون النص في كل منها هو `list item`. يجب أن تتطابق الحروف الكبيرة والفواصل تمامًا.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

يجب أن يصل شفرتك إلى الكائن `d3`.

```js
assert.match(code, /d3/g);
```

يجب أن تستخدم شفرتك الطريقة `selectAll`.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
