---
id: 587d7fa8367417b2b2512bcd
title: إنشاء شريط لكل نقطة بيانات في المجموعة
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

أضاف التحدٍّ الأخير مستطيلاً واحدًا فقط إلى عنصر `svg` لتمثيل شريط. هنا، ستدمج ما تعلمته حتى الآن عن `data()` و`enter()` وأشكال SVG لإنشاء مستطيل وإضافته لكل نقطة بيانات في `dataset`.

أظهر تحدٍّ سابق صيغة كيفية إنشاء وإضافة عنصر `div` لكل عنصر في `dataset`:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

هناك بعض الاختلافات عند العمل مع عناصر `rect` بدلاً من عناصر `div`. يجب إضافة عناصر `rect` إلى عنصر `svg`، وليس مباشرة إلى `body`. أيضًا، تحتاج إلى إخبار D3 بمكان وضع كل `rect` داخل مساحة `svg`. سيتم تغطية وضع الأشرطة في التحدٍّ التالي.

# --instructions--

استخدم الطرق `data()` و`enter()` و`append()` لإنشاء وإضافة عنصر `rect` لكل عنصر في `dataset`. يجب أن تظهر الأشرطة كلها فوق بعضها البعض؛ وسيتم إصلاح ذلك في التحدٍّ التالي.

# --hints--

يجب أن يحتوي مستندك على 9 عناصر `rect`.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

يجب أن يستخدم شفرتك الطريقة `data()`.

```js
assert.match(code, /\.data/g);
```

يجب أن يستخدم شفرتك الطريقة `enter()`.

```js
assert.match(code, /\.enter/g);
```

يجب أن يستخدم شفرتك الطريقة `append()`.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
