---
id: 587d7fa7367417b2b2512bc5
title: العمل مع البيانات الديناميكية في D3
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

التحديان الأخيران يغطيان أساسيات عرض البيانات ديناميكيًا باستخدام D3 عبر الطريقتين `data()` و `enter()`. هذه الطرق تأخذ مجموعة بيانات وبالاشتراك مع طريقة `append()` تُنشئ عنصر DOM جديد لكل مدخل في مجموعة البيانات.

في التحدي السابق، أنشأت عنصر `h2` جديد لكل عنصر في مصفوفة `dataset`، لكن جميعها احتوت على نفس النص، `New Title`. السبب هو أنك لم تستخدم البيانات المرتبطة بكل عنصر من عناصر `h2`.

طريقة D3 `text()` يمكن أن تأخذ سلسلة نصية أو دالة استدعاء راجع كمعلمة:

```js
selection.text(d => d);
```

في المثال أعلاه، المعلمة `d` تشير إلى مدخل واحد في مجموعة البيانات المرتبطة بالاختيار.

باستخدام المثال الحالي كمرجع، العنصر الأول `h2` مرتبط بالقيمة 12، والعنصر الثاني `h2` مرتبط بالقيمة 31، والعنصر الثالث `h2` مرتبط بالقيمة 22، وهكذا.

# --instructions--

غيّر طريقة `text()` بحيث يعرض كل عنصر `h2` القيمة المقابلة من مصفوفة `dataset` مع مسافة واحدة وسلسلة النص `USD`. على سبيل المثال، يجب أن يكون العنوان الأول `12 USD`.

# --hints--

يجب أن يحتوي العنصر الأول `h2` على النص `12 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

يجب أن يحتوي العنصر الثاني `h2` على النص `31 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

يجب أن يحتوي العنصر الثالث `h2` على النص `22 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

يجب أن يحتوي العنصر الرابع `h2` على النص `17 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

يجب أن يحتوي العنصر الخامس `h2` على النص `25 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

يجب أن يحتوي العنصر السادس `h2` على النص `18 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

يجب أن يحتوي العنصر السابع `h2` على النص `29 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

يجب أن يحتوي العنصر الثامن `h2` على النص `14 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

يجب أن يحتوي العنصر التاسع `h2` على النص `9 USD`.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
