---
id: 587d7fa7367417b2b2512bc4
title: العمل مع البيانات في D3
challengeType: 6
forumTopicId: 301497
dashedName: work-with-data-in-d3
---

# --description--

تركز مكتبة D3 على النهج القائم على البيانات. عندما يكون لديك مجموعة من البيانات، يمكنك تطبيق طرق D3 لعرضها على الصفحة. تأتي البيانات بأشكال متعددة، لكن هذا التحدٍّ يستخدم مصفوفة بسيطة من الأرقام.

الخطوة الأولى هي تعريف D3 بالبيانات. تُستخدم الطريقة `data()` على اختيار لعناصر DOM لربط البيانات بتلك العناصر. تُمرر مجموعة البيانات كمعلمة للطريقة.

نمط سير العمل الشائع هو إنشاء عنصر جديد في المستند لكل قطعة بيانات في المجموعة. لدى D3 الطريقة `enter()` لهذا الغرض.

عندما تُدمج `enter()` مع الطريقة `data()`، تنظر إلى العناصر المختارة من الصفحة وتقارنها بعدد عناصر البيانات في المجموعة. إذا كان عدد العناصر أقل من عدد عناصر البيانات، تُنشئ العناصر المفقودة.

إليك مثالًا يختار عنصر `ul` وينشئ عنصر قائمة جديد بناءً على عدد المدخلات في المصفوفة:

```html
<body>
  <ul></ul>
  <script>
    const dataset = ['a', 'b', 'c'];
    d3.select('ul')
      .selectAll('li')
      .data(dataset)
      .enter()
      .append('li')
      .text('New item');
  </script>
</body>
```

قد يبدو اختيار عناصر غير موجودة بعد أمرًا محيرًا. هذا الشفرة تخبر D3 أولًا باختيار `ul` على الصفحة. بعد ذلك، اختيار كل عناصر القائمة، مما يرجع اختيارًا فارغًا. ثم تراجع الطريقة `data()` مجموعة البيانات وتنفذ الشفرة التالية ثلاث مرات، مرة لكل عنصر في المصفوفة. ترى الطريقة `enter()` أنه لا توجد عناصر `li` على الصفحة، لكنها تحتاج إلى 3 (واحدة لكل قطعة بيانات في `dataset`). تُضاف عناصر `li` جديدة إلى `ul` وتحمل النص `New item`.

# --instructions--

اختر عقدة `body`، ثم اختر كل عناصر `h2`. اجعل D3 ينشئ ويضيف وسم `h2` لكل عنصر في مصفوفة `dataset`. يجب أن يكون النص في `h2` هو `New Title`. يجب أن يستخدم شفرتك الطريقتين `data()` و `enter()`.

# --hints--

يجب أن يحتوي مستندك على 9 عناصر `h2`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 9);
```

يجب أن يكون النص في عناصر `h2` هو `New Title`. يجب أن تتطابق الحروف الكبيرة والمسافات تمامًا.

```js
const h2Elements = document.querySelectorAll('h2');
for (let i = 0; i < h2Elements.length; i++) {
  assert.match(h2Elements[i]?.textContent, /New Title/g);
}
```

يجب أن تستخدم شفرتك الطريقة `data()`.

```js
assert.match(code, /\.data/g);
```

يجب أن تستخدم شفرتك الطريقة `enter()`.

```js
assert.match(code, /\.enter/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text('New Title');
  </script>
</body>
```
