---
id: 587d7faa367417b2b2512bd6
title: إضافة تلميح إلى عنصر D3
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

يعرض التلميح معلومات إضافية عن عنصر في الصفحة عندما يحوم المستخدم فوق ذلك العنصر. هناك عدة طرق لإضافة تلميح إلى التصور. يستخدم هذا التحدٍّ عنصر SVG `title`.

يرتبط `title` بطريقة `text()` لإضافة البيانات ديناميكيًا إلى الأعمدة.

# --instructions--

أضف عنصر `title` تحت كل عقدة `rect`. ثم نفذ طريقة `text()` مع دالة استدعاء راجع ليعرض النص قيمة البيانات.

# --hints--

يجب أن يحتوي شفرتك على 9 عناصر `title`.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

يجب أن يحتوي أول عنصر `title` على نص التلميح `12`.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

يجب أن يحتوي ثاني عنصر `title` على نص التلميح `31`.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

يجب أن يحتوي ثالث عنصر `title` على نص التلميح `22`.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

يجب أن يحتوي رابع عنصر `title` على نص التلميح `17`.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

يجب أن يحتوي خامس عنصر `title` على نص التلميح `25`.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

يجب أن يحتوي سادس عنصر `title` على نص التلميح `18`.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

يجب أن يحتوي سابع عنصر `title` على نص التلميح `29`.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

يجب أن يحتوي ثامن عنصر `title` على نص التلميح `14`.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

يجب أن يحتوي تاسع عنصر `title` على نص التلميح `9`.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
