---
id: 587d7fa9367417b2b2512bd0
title: عكس عناصر SVG
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

ربما لاحظت أن الثميل البياني بالأعمدة يبدو مقلوبًا أو معكوسًا. هذا بسبب طريقة استخدام SVG لإحداثيات (x, y).

في SVG، نقطة الأصل للإحداثيات تكون في الزاوية العلوية اليسرى. إحداثي `x` بقيمة 0 يضع الشكل على الحافة اليسرى لمساحة SVG. إحداثي `y` بقيمة 0 يضع الشكل على الحافة العلوية لمساحة SVG. القيم الأعلى لـ `x` تدفع المستطيل إلى اليمين. القيم الأعلى لـ `y` تدفع المستطيل إلى الأسفل.

لجعل الأعمدة في وضعها الصحيح، تحتاج إلى تغيير طريقة حساب إحداثي `y`. يجب أن يأخذ في الاعتبار كل من ارتفاع العمود والارتفاع الكلي لمساحة SVG.

ارتفاع مساحة SVG هو 100. إذا كان لديك نقطة بيانات بقيمة 0 في المجموعة، فستريد أن يبدأ العمود من أسفل مساحة SVG (وليس من الأعلى). لتحقيق ذلك، يجب أن يكون لإحداثي `y` قيمة 100. إذا كانت قيمة نقطة البيانات 1، تبدأ بإحداثي `y` بقيمة 100 لوضع العمود في الأسفل. ثم تحتاج إلى احتساب ارتفاع العمود الذي يساوي 1، لذا سيكون الإحداثي النهائي لـ `y` هو 99.

الإحداثي `y` الذي هو `y = heightOfSVG - heightOfBar` سيضع الأعمدة في وضعها الصحيح.

# --instructions--

غيّر دالة الاستدعاء الراجع لخاصية `y` لوضع الأعمدة في وضعها الصحيح. تذكر أن `height` للعمود هو 3 أضعاف قيمة البيانات `d`.

**ملاحظة:** بشكل عام، العلاقة هي `y = h - m * d`، حيث `m` هو الثابت الذي يضبط مقياس نقاط البيانات.

# --hints--

يجب أن يكون للإطار الأول `rect` قيمة `y` تساوي `64`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

يجب أن يكون للإطار الثاني `rect` قيمة `y` تساوي `7`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

يجب أن يكون للإطار الثالث `rect` قيمة `y` تساوي `34`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

يجب أن يكون للإطار الرابع `rect` قيمة `y` تساوي `49`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

يجب أن يكون للإطار الخامس `rect` قيمة `y` تساوي `25`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

يجب أن يكون للإطار السادس `rect` قيمة `y` تساوي `46`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

يجب أن يكون للإطار السابع `rect` قيمة `y` تساوي `13`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

يجب أن يكون للإطار الثامن `rect` قيمة `y` تساوي `58`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

يجب أن يكون للإطار التاسع `rect` قيمة `y` تساوي `73`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
