---
id: 587d7fab367417b2b2512bd9
title: إضافة مسميات إلى دوائر المخطط الانتشاري
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

يمكنك إضافة نصوص لإنشاء مسميات للنقاط في المخطط الانتشاري.

الهدف هو عرض القيم المفصولة بفاصلة للحقل الأول (`x`) والثاني (`y`) لكل عنصر في `dataset`.

تحتاج عقد `text` إلى خاصيتي `x` و`y` لتحديد موقعها على SVG. في هذا التحدٍّ، يمكن لقيمة `y` (التي تحدد الارتفاع) أن تستخدم نفس القيمة التي يستخدمها `circle` في خاصية `cy`. يمكن أن تكون قيمة `x` أكبر قليلاً من قيمة `cx` في `circle`، ليكون المسمى مرئيًا. هذا سيدفع المسمى إلى يمين النقطة المرسومة.

# --instructions--

علّم كل نقطة على مخطط التبعثر باستخدام عناصر `text`. يجب أن يكون نص المسمى هو القيمتان مفصولتان بفاصلة ومسافة. على سبيل المثال، المسمى للنقطة الأولى هو `34, 78`. عيّن الخاصية `x` بحيث تكون أكبر بـ `5` وحدات من القيمة التي استخدمتها في الخاصية `cx` على العنصر `circle`. عيّن الخاصية `y` بنفس الطريقة التي تُستخدم للقيمة `cy` على العنصر `circle`.

# --hints--

يجب أن يحتوي شفرتك على 10 عناصر `text`.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

يجب أن يحتوي المسمى الأول على نص `34, 78`، وقيمة `x` تساوي `39`، وقيمة `y` تساوي `422`.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

يجب أن يحتوي المسمى الثاني على نص `109, 280`، وقيمة `x` تساوي `114`، وقيمة `y` تساوي `220`.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

يجب أن يحتوي المسمى الثالث على نص `310, 120`، وقيمة `x` تساوي `315`، وقيمة `y` تساوي `380`.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

يجب أن يحتوي المسمى الرابع على نص `79, 411`، وقيمة `x` تساوي `84`، وقيمة `y` تساوي `89`.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

يجب أن يحتوي المسمى الخامس على نص `420, 220`، وقيمة `x` تساوي `425`، وقيمة `y` تساوي `280`.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

يجب أن يحتوي المسمى السادس على نص `233, 145`، وقيمة `x` تساوي `238`، وقيمة `y` تساوي `355`.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

يجب أن يحتوي المسمى السابع على نص `333, 96`، وقيمة `x` تساوي `338`، وقيمة `y` تساوي `404`.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

يجب أن يحتوي المسمى الثامن على نص `222, 333`، وقيمة `x` تساوي `227`، وقيمة `y` تساوي `167`.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

يجب أن يحتوي المسمى التاسع على نص `78, 320`، وقيمة `x` تساوي `83`، وقيمة `y` تساوي `180`.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

يجب أن يحتوي المسمى العاشر على نص `21, 123`، وقيمة `x` تساوي `26`، وقيمة `y` تساوي `377`.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
