---
id: 587d7fac367417b2b2512bdc
title: استخدام دالتي d3.max و d3.min لإيجاد القيم الحدية الدنيا والعليا في مجموعة بيانات
  Use the d3.max and d3.min Functions to Find Minimum and Maximum Values in a
  Dataset
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

تُحدد دالتي D3 `domain()` و `range()` تلك المعلومات لمقياسك بناءً على البيانات. هناك طريقتان لتسهيل ذلك.

غالبًا عند تعيين المجال، سترغب في استخدام القيم الحدية الدنيا والعليا ضمن مجموعة البيانات. محاولة إيجاد هذه القيم يدويًا، خاصة في مجموعة بيانات كبيرة، قد تسبب أخطاء.

تحتوي D3 على دالتين - `min()` و `max()` لإرجاع هذه المعلومات. إليك مثالًا:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

قد تحتوي مجموعة البيانات على مصفوفات مسكنة، مثل أزواج الإحداثيات `[x, y]` التي كانت في مثال مخطط التبعثر. في هذه الحالة، تحتاج إلى إخبار D3 بكيفية حساب الحد الأدنى والأقصى. لحسن الحظ، تأخذ كل من دالتي `min()` و `max()` دالة استدعاء راجع. في هذا المثال، المعلمة `d` في دالة الاستدعاء الراجع تمثل المصفوفة الداخلية الحالية. يجب أن تُرجع دالة الاستدعاء الراجع العنصر من المصفوفة الداخلية (قيمة `x` أو `y`) التي تريد حساب الحد الأدنى أو الأقصى عليها. إليك مثالًا على كيفية إيجاد القيم الحدية الدنيا والعليا باستخدام مصفوفة من المصفوفات:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

سيكون لـ `minX` القيمة `1`.

# --instructions--

تحتوي مصفوفة `positionData` على مصفوفات فرعية لإحداثيات x و y و z. استخدم دالة من D3 لإيجاد القيمة العظمى لإحداثي z (القيمة الثالثة) من المصفوفات واحفظها في المتغير `output`.

# --hints--

يجب أن يكون النص في `h2` هو `8`.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

يجب أن يستخدم شفرتك الدالة `max()`.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
