---
id: 587d7fa8367417b2b2512bcb
title: تعلّم عن SVG في D3
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> تعني <dfn>الرسومات المتجهة القابلة للتوسع</dfn>.

هنا تعني كلمة "قابلة للتوسع" أنه إذا قمت بالتكبير أو التصغير على عنصر، فلن يظهر بشكل متكسّر. يتوسع مع نظام العرض، سواء كان على شاشة هاتف صغيرة أو شاشة تلفاز كبيرة.

يُستخدم SVG لإنشاء الأشكال الهندسية الشائعة. بما أن D3 يحول البيانات إلى تمثيل بصري، فإنه يستخدم SVG لإنشاء الأشكال الخاصة بالتمثيل. يجب أن توضع أشكال SVG الخاصة بصفحة ويب داخل وسم HTML `svg`.

يمكن أن يكون CSS قابلاً للتوسع عندما تستخدم الأنماط وحدات نسبية (مثل `vh`، `vw`، أو النسب المئوية)، لكن استخدام SVG أكثر مرونة لبناء تمثيلات البيانات.

# --instructions--

أضف عقدة `svg` إلى `body` باستخدام `append()`. أعطها خاصية `width` بقيمة الثابت `w` المعطى وخاصية `height` بقيمة الثابت `h` المعطى باستخدام طريقتي `attr()` أو `style()` لكل منهما. سترى ذلك في الناتج لأن هناك `background-color` باللون الوردي مطبق عليها في وسم `style`.

**ملاحظة:** عند استخدام `attr()` لا تحتوي خاصيتي العرض والارتفاع على وحدات. هذا هو الأساس في التوسع - العنصر سيحافظ دائمًا على نسبة عرض إلى ارتفاع 5:1، بغض النظر عن مستوى التكبير.

# --hints--

يجب أن يحتوي مستندك على عنصر `svg` واحد.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

يجب أن يحتوي عنصر `svg` على خاصية `width` مضبوطة على `500` أو منسَّقة لتكون بعرض `500px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

يجب أن يحتوي عنصر `svg` على خاصية `height` مضبوطة على `100` أو منسَّقة لتكون بارتفاع `100px`.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
