---
id: 587d7fa9367417b2b2512bd1
title: تغيير لون عنصر SVG
challengeType: 6
forumTopicId: 301480
dashedName: change-the-color-of-an-svg-element
---

# --description--

الأشرطة في الموضع الصحيح، لكنها كلها بنفس اللون الأسود. لدى SVG طريقة لتغيير لون الأشرطة.

في SVG، يُلوَّن شكل `rect` باستخدام الخاصية `fill`. تدعم رموز الألوان السداسية، وأسماء الألوان، وقيم rgb، بالإضافة إلى خيارات أكثر تعقيدًا مثل التدرجات والشفافية.

# --instructions--

أضف طريقة `attr()` لتعيين الخاصية `fill` لكل الأشرطة إلى اللون البحري.

# --hints--

يجب أن يكون لكل الأشرطة لون `fill` باللون البحري.

```js
const rect = document.querySelector('rect');
assert.exists(rect); 
const fill = window.getComputedStyle(rect)['fill'];
assert.strictEqual(fill, 'rgb(0, 0, 128)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
  </script>
</body>
```
