---
id: 587d7fa9367417b2b2512bce
title: تعيين الإحداثيات ديناميكيًا لكل عمود
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

التحدي الأخير أنشأ وأضاف مستطيلاً إلى عنصر `svg` لكل نقطة في `dataset` لتمثيل عمود. للأسف، كانت جميعها مكدسة فوق بعضها البعض.

يُحدد موضع المستطيل بواسطة الخاصيتين `x` و `y`. تخبران D3 أين تبدأ رسم الشكل في مساحة `svg`. التحدي الأخير جعل قيمهما 0، لذا وُضع كل عمود في الزاوية العلوية اليسرى.

في المخطط البياني بالأعمدة، يجب أن تكون كل الأعمدة على نفس المستوى الرأسي، مما يعني أن قيمة `y` تبقى ثابتة (عند 0) لكل الأعمدة. أما قيمة `x` فيجب أن تتغير مع إضافة أعمدة جديدة. تذكر أن القيم الأكبر لـ `x` تدفع العناصر إلى اليمين أكثر. أثناء التكرار على عناصر المصفوفة في `dataset`، يجب أن تزداد قيمة `x`.

تقبل طريقة `attr()` في D3 دالة استدعاء راجع لتعيين الخاصية ديناميكيًا. تأخذ دالة الاستدعاء راجع معلمتين، واحدة لنقطة البيانات نفسها (عادة `d`) وواحدة لفهرس نقطة البيانات في المصفوفة. المعلمة الثانية للفهرس اختيارية. هذا هو التنسيق:

```js
selection.attr('property', (d, i) => {});
```

من المهم أن تعرف أنك لا تحتاج إلى كتابة حلقة `for` أو استخدام `forEach()` للتكرار على عناصر مجموعة البيانات. تذكر أن طريقة `data()` تُحلل مجموعة البيانات، وأي طريقة تُسلسل بعد `data()` تُنفذ مرة واحدة لكل عنصر في مجموعة البيانات.

# --instructions--

غيّر دالة الاستدعاء راجع للخاصية `x` بحيث تُرجع الفهرس مضروبًا في 30.

**ملاحظة:** عرض كل عمود هو 25، لذا زيادة كل قيمة `x` بمقدار 30 تضيف بعض المسافة بين الأعمدة. أي قيمة أكبر من 25 ستعمل في هذا المثال.

# --hints--

يجب أن يكون للعنصر `rect` الأول قيمة `x` تساوي `0`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

يجب أن يكون للعنصر `rect` الثاني قيمة `x` تساوي `30`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

يجب أن يكون للعنصر `rect` الثالث قيمة `x` تساوي `60`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

يجب أن يكون للعنصر `rect` الرابع قيمة `x` تساوي `90`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

يجب أن يكون للعنصر `rect` الخامس قيمة `x` تساوي `120`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

يجب أن يكون للعنصر `rect` السادس قيمة `x` تساوي `150`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

يجب أن يكون للعنصر `rect` السابع قيمة `x` تساوي `180`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

يجب أن يكون للعنصر `rect` الثامن قيمة `x` تساوي `210`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

يجب أن يكون للعنصر `rect` التاسع قيمة `x` تساوي `240`.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
