---
id: 68e4c81577243e4db1676491
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次研讨会中，你将通过创建一个包含 Quincy Larson 求职开发者职位技巧的网页来练习使用语义化 `HTML`。基本的 `HTML` 公式化已为你准备好。

首先创建一个带有文本 `Quincy's Tips for Getting a Developer Job` 的 `h1` 元素。

# --hints--

你应该有一个 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素应包含文本 `Quincy's Tips for Getting a Developer Job`。

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
