---
id: 68e65ec5348e40a739903e6d
title: 步骤2
challengeType: 0
dashedName: step-2
---

# --description--

现在，在 `h1` 元素下创建一个段落元素。在此段落内添加文本 `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

你应该有一个 `p` 元素。

```js
assert.exists(document.querySelector('p'));
```

你的段落应包含文本 `Learning to code is hard, but as Quincy Larson says, You can become a developer.`。请仔细查看你的拼写。

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

你的段落应位于 `h1` 元素下方。

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
