---
id: 68e964b84f954c322376e592
title: 第 4 步
challengeType: 0
dashedName: step-4
---

# --description--

如果引用的来源是网站，你可以使用 `cite` 属性来引用它。该属性的值应是一个有效的 URL。虽然这个属性不会改变引用块的显示形式，但它对为屏幕阅读器和搜索引擎提供更多关于引用的信息非常有帮助。

这是一个带有 `cite` 属性的内联引用元素示例：

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

为你的内联引用元素添加带有此 URL 的 `cite` 属性：

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

你的 `q` 元素应具有 `cite` 属性。

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

你的 `q` 元素的 `cite` 属性应具有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
