---
id: 68e90335bd53a901c41dfc13
title: 步骤 3
challengeType: 0
dashedName: step-3
---

# --description--

在之前的课程中，你学习了 `q` 元素用于区分引用文本与周围内容。

这是一个示例：

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

大多数现代浏览器在你使用 `q` 元素时会自动为内联引用添加引号。

在段落内将 `You can become a developer.` 包裹在内联引用元素中，其余段落保持不变。

# --hints--

你应该有一个 `q` 元素。

```js
assert.exists(document.querySelector('q'));
```

你的 `q` 元素应嵌套在位于 `h1` 元素下方的段落内。

```js
assert.exists(document.querySelector('h1 + p > q'));
```

你的 `q` 元素应包含文本 `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

在你的 `q` 元素内，文本不得有前导或尾随空白（第一个字符之前或最后一个字符之后不能有空格、标签（页）或换行符）。

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

你应该只在句子的适当部分添加 `q` 标签，而不更改原始文本。`p` 元素仍应显示：`Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
