---
id: 68e964b84f954c322376e592
title: 步骤4
challengeType: 0
dashedName: step-4
---

# --description--

If the source of a quote is a website, you can cite it with the `cite` attribute. 该属性的值应为有效的 URL。 虽然该属性不会改变块级引用的显示效果，但它非常有助于为屏幕阅读器和搜索引擎提供有关引文的更多信息。

这是一个带有 `cite` 属性的内联引用元素示例：

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

为你的内联引用元素添加带有此 URL 的 `cite` 属性：

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

你的 `q` 元素应具有 `cite` 属性。

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

你的 `q` 元素的 `cite` 属性应具有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
