---
id: 68e96ce52087043b0e999296
title: 步骤5
challengeType: 0
dashedName: step-5
---

# --description--

在段落元素下方，添加一个 `main` 元素，并在其中嵌套三个 `section` 元素。

# --hints--

你应该在段落元素下方有一个 `main` 元素。

```js
const mainEl = document.querySelector('p + main');
assert.exists(mainEl);
```

你的 `main` 元素内应包含三个 `section` 元素。

```js
const sections = document.querySelectorAll('main > section');
assert.lengthOf(sections, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
