---
id: 68ea5ccc1229b85eca38a3e8
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在第一个 `section` 元素内，添加一个带有文本 `Envisioning Success` 的 `h2` 元素。

# --hints--

在第一个 `section` 元素内，你应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector('section:first-of-type > h2'));
```

你的 `h2` 应该有文本 `Envisioning Success`。

```js
const firstH2El = document.querySelector('section:first-of-type > h2');
assert.equal(firstH2El?.innerText.trim(), `Envisioning Success`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
