---
id: 68ea7e7b916e8270483689b6
title: 步骤7
challengeType: 0
dashedName: step-7
---

# --description--

在 `h2` 元素下方，你将添加另一个来自 Quincy 的引用。这次，引用不会是更大段落的一部分。相反，整个段落将是一个引用。为了区分这样的引用文本，你应该使用块引用元素：`blockquote`。在浏览器中，你会看到文本稍微缩排。

这是一个带有引用文本的块引用元素示例：

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

现在，在第一个部分内，在 `h2` 元素下方添加一个块引用元素，文本为 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`。

# --hints--

在第一个 `section` 元素内，你应该在 `h2` 元素下方有一个 `blockquote` 元素。

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

你的 `blockquote` 元素应包含文本 `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`。

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
