---
id: 68ebe9cd7523d340c0b343d3
title: 步骤 13
challengeType: 0
dashedName: step-13
---

# --description--

在第三个也是最后一个 `section` 元素中，你将创建与前述的 `section` 中相同的超文本标记语言结构。

首先，在最后一节内，添加一个带有文本 `Importance of Building a Reputation` 的 `h2` 元素。

然后，在 `h2` 元素下方，添加一个块引用元素，其 `cite` 属性设置为 `https://www.freecodecamp.org/news/learn-to-code-book/`。

最后，在块引用元素内依次嵌套三个段落元素，文本内容如下：

1. `Share short video demos of your projects.`
2. `Keep applying to speak at bigger and bigger conferences.`
3. `Hang out at hackerspaces and help people who are even newer to coding than you.`

恭喜！你已经完成了本次研讨会。

# --hints--

你应该在第三个部分内嵌套一个 `h2` 元素。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2'));
```

你的 `h2` 元素应包含文本 `Importance of Building a Reputation`。

```js
const h2El = document.querySelector('main > section:nth-of-type(3) > h2');
assert.equal(h2El?.innerText.trim(), 'Importance of Building a Reputation');
```

你应该在第三个部分内嵌套一个 `blockquote` 元素。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > blockquote'));
```

你的 `blockquote` 元素应位于你的 `h2` 元素下方。

```js
assert.exists(document.querySelector('main > section:nth-of-type(3) > h2 + blockquote'));
```

你的 `blockquote` 元素应具有 `cite` 属性。

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.exists(blockquoteEl?.getAttribute('cite'));
```

你的 `blockquote` 元素的 `cite` 属性应具有值 `https://www.freecodecamp.org/news/learn-to-code-book/`。

```js
const blockquoteEl = document.querySelector('main > section:nth-of-type(3) > blockquote');
assert.equal(blockquoteEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

你应该在第三部分的 `blockquote` 元素内嵌套三个 `p` 元素。

```js
const pElements = document.querySelectorAll('main > section:nth-of-type(3) > blockquote > p');
assert.lengthOf(pElements, 3);
```

你的第一个 `p` 元素应包含文本 `Share short video demos of your projects.`

```js
const firstPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(1)');
assert.equal(firstPEl?.innerText.trim(), 'Share short video demos of your projects.');
```

你的第二个 `p` 元素应包含文本 `Keep applying to speak at bigger and bigger conferences.`

```js
const secondPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(2)');
assert.equal(secondPEl?.innerText.trim(), "Keep applying to speak at bigger and bigger conferences.");
```

你的第三个 `p` 元素应包含文本 `Hang out at hackerspaces and help people who are even newer to coding than you.`

```js
const thirdPEl = document.querySelector('main > section:nth-of-type(3) > blockquote > p:nth-of-type(3)');
assert.equal(thirdPEl?.innerText.trim(), 'Hang out at hackerspaces and help people who are even newer to coding than you.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>
      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>
      <section>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
          &mdash;Quincy Larson, <cite>How to Learn to Code and Get a Developer Job [Full Book]</cite>
        </p>
      </section>

      <section>
        <h2>Importance of Networking</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>So much of getting a job is who you know.</p>
          <p>It's OK to be an introvert, but you do need to push your boundaries.</p>
          <p>Create GitHub, Twitter, LinkedIn, and Discord accounts.</p>
          <p>Go to tech meetups and conferences. Travel if you have to.</p>
        </blockquote>
      </section>

      <section>
        <h2>Importance of Building a Reputation</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          <p>Share short video demos of your projects.</p>
          <p>Keep applying to speak at bigger and bigger conferences.</p>
          <p>Hang out at hackerspaces and help people who are even newer to coding than you.</p>
        </blockquote>
      </section>
    </main>
  </body>
</html>
```
