---
id: 69a00f673f171e29159e3ad8
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Esse nó precisaria de duas propriedades, o elemento a ser adicionado e um ponteiro para o próximo nó.

Dentro da função add, crie um objeto chamado `node`, ele deve ter uma propriedade chamada `element` e uma propriedade `next` definida como `null`, que é o ponteiro para o próximo elemento da lista.

# --hints--

Você deve criar um objeto chamado `node` dentro da sua função `add`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /const|var|let\s+node\s*=\s*{/);
```

Seu objeto `node` deve ter uma propriedade chamada `element` definida como `element`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /element\s*:\s*element/);
```

Seu objeto `node` deve ter uma propriedade chamada `next` definida como `null`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /next\s*:\s*null/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}
```
