---
id: 69a0179a85b00b6127e0af45
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Dentro do loop `while`, atualize `current` para `current.next` para que o loop avance pela lista um nó por vez, e pare quando alcançar o nó com uma propriedade `next` igual a `null`, ou seja, o fim da lista.

# --hints--

Você deve atualizar o valor da sua variável `current` para `current.next` dentro do loop `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            --fcc-editable-region--

            --fcc-editable-region--
        }
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
