---
id: 69a03aaa1e6fe8152788ce6a
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Dentro da função `remove`, use `let` para criar uma variável `previous` definida como `null`, e `let` novamente para criar uma variável `current` definida como `list.head`. Essas são as variáveis que você usará no loop que percorrerá a lista.

`current` acompanha o nó que está sendo verificado na lista, começando pela cabeça e avançando um nó por vez. `previous`, por outro lado, mantém o controle do nó logo antes do atual.

Isso é importante ao remover um nó porque você deve atualizar o valor next do nó anterior para que ele aponte para o nó depois daquele que está sendo removido, para manter a lista conectada.

# --hints--

Você deve usar `let` para criar uma variável `previous` definida como `null` dentro da sua função `remove`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const { previous } = remove?.variables ?? {};
assert.isTrue(previous?.toString().startsWith('let'));
assert.isTrue(previous?.value.matches('null'));
```

Você deve usar `let` para criar uma variável `current` definida como `list.head` dentro da sua função `remove`.

```js
const explorer = await __helpers.Explorer(code);
const { remove } = explorer.allFunctions;
const { current } = remove?.variables ?? {};
assert.isTrue(current?.toString().startsWith('let'));
assert.isTrue(current?.value.matches('list.head'));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList);
console.log(isEmpty(myList));
```
