---
id: 69a00a051b9a18694d00c16b
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, vocĂȘ vai implementar uma lista ligada.

Para comeĂ§ar, crie uma funĂ§ĂŁo chamada `initList`. Dentro dela, retorne um objeto com uma propriedade `head` definida como `null` e uma propriedade `length` com valor `0`.

Esta Ă© a funĂ§ĂŁo que vocĂȘ vai chamar para inicializar uma lista ligada. O `head` representa o primeiro elemento da lista, inicialmente, ele Ă© `null` porque nĂŁo hĂĄ nenhum elemento na lista. O `length` acompanha o nĂșmero de nĂłs (elementos) na lista, que vocĂȘ vai incrementar ou decrementar ao adicionar ou remover um nĂł da lista.

# --hints--

VocĂȘ deve criar uma funĂ§ĂŁo `initList`.

```js
assert.isFunction(initList)
```

Sua funĂ§ĂŁo `initList` deve retornar um objeto.

```js
assert.isObject(initList());
```

Sua funĂ§ĂŁo `initList` deve retornar um objeto com um `head` definido como `null`.

```js
assert.propertyVal(initList(), "head", null);
```

Sua funĂ§ĂŁo `initList` deve retornar um objeto com um `length` definido como `0`.

```js
assert.propertyVal(initList(), "length", 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
