---
id: 69a03aaa1e6fe8152788ce6a
title: passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Dentro da funĂ§ĂŁo `remove`, use `let` para criar uma variĂĄvel `previous` definida como `null`, e `let` novamente para criar uma variĂĄvel `current` definida como `list.head`. Essas sĂŁo as variĂĄveis que vocĂȘ usarĂĄ no loop que percorrerĂĄ a lista.

`current` acompanha o nĂł que estĂĄ sendo verificado na lista, comeĂ§ando pela cabeĂ§a e avanĂ§ando um nĂł por vez. `previous`, por outro lado, mantĂ©m o controle do nĂł logo antes do atual.

Isso Ă© importante ao remover um nĂł porque vocĂȘ deve atualizar o valor next do nĂł anterior para que ele aponte para o nĂł depois daquele que estĂĄ sendo removido, para manter a lista conectada.

# --hints--

VocĂȘ deve usar `let` para criar uma variĂĄvel `previous` definida como `null` dentro da sua funĂ§ĂŁo `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+previous\s*=\s*null;?/);
```

VocĂȘ deve usar `let` para criar uma variĂĄvel `current` definida como `list.head` dentro da sua funĂ§ĂŁo `remove`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    --fcc-editable-region--
    
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
