---
id: 69a00eaa43458e618331873c
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Agora você vai trabalhar na função que vai adicionar um nó à lista.

Crie uma função `add` com os parâmetros `list` e `element`.

# --hints--

Você deve criar uma função `add`.

```js
assert.isFunction(add)
```

Sua função `add` deve ter os parâmetros `list` e `element`.

```js
const addParams = __helpers.getFunctionParams(add.toString());
const addParamsNames = addParams.map(param => param.name)
assert.include(addParamsNames, "list");
assert.include(addParamsNames, "element");
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

--fcc-editable-region--

--fcc-editable-region--
```
