---
id: 69a013a86c97dac3784b2636
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora, o log deve mostrar algo como `{ head: { element: 42, next: null }, length: 1 }` no console.

No entanto, essa não é a forma de lidar com um caso em que o nó adicionado é o primeiro elemento da lista, ou seja, a cabeça.

Para fazer isso adequadamente, use uma instrução if para verificar se a lista está vazia chamando a função `isEmpty` com o `list`. Se isso for verdadeiro, então mova sua linha `list.head = node` para o corpo da declaração `if`.

# --hints--

Você deve ter uma declaração `if` com a condição `isEmpty(list)`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){');
```

Você deve mover sua linha `list.head = node` para dentro da declaração `if`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'if(isEmpty(list)){list.head=node');
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

--fcc-editable-region--
    list.head = node;
--fcc-editable-region--
    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
