---
id: 69a045fd236d0d973a59f103
title: Passo 19
challengeType: 1
dashedName: step-19
---

# --description--

Se `previous` é `null`, isso significa que o elemento a ser removido está no início da lista, então você deve atualizar o `head` para apontar para o próximo nó para remover o primeiro elemento.

Para fazer isso, adicione uma cláusula `else`. Dentro dela, defina `list.head` como `current.next`.

# --hints--

Você deve adicionar uma cláusula `else` à sua declaração `if`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*/);
```

Você deve definir `list.head` como `current.next` dentro da cláusula `else`.

```js
assert.match(__helpers.removeJSComments(remove.toString().split("previous.next = current.next;")), /\}\s*else\s*\{\s*list\.head\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }

    if(previous !== null) {
        previous.next = current.next;
    --fcc-editable-region--
    } 
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
