---
id: 69a00cd4ec513345d40e95ec
title: passo 2
challengeType: 1
dashedName: step-2
---

# --description--

A próxima coisa em que você vai trabalhar é uma função que permite verificar se a lista está vazia.

Para isso, crie uma função chamada `isEmpty` que recebe um parâmetro `list`. Dentro da função, retorne uma verificação para ver se a propriedade `length` de `list` é estritamente igual a `0`.

# --hints--

Você deve criar uma função `isEmpty`.

```js
assert.isFunction(isEmpty)
```

Sua função `isEmpty` deve ter um parâmetro `list`.

```js
const isEmptyParams = __helpers.getFunctionParams(isEmpty.toString());
assert.deepEqual(isEmptyParams[0].name, "list");
```

Sua função `isEmpty` deve retornar `list.length === 0`.

```js
const emptyList = { head: null, length: 0 };
const nonEmptyList = { head: {}, length: 1 };

assert.isTrue(isEmpty(emptyList));
assert.isFalse(isEmpty(nonEmptyList));
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
