---
id: 69a04466917c66b132323a1b
title: Passo 18
challengeType: 1
dashedName: step-18
---

# --description--

Se `previous` não for `null`, isso significa que o elemento a ser removido não está no início, mas em algum lugar dentro da lista, então o nó atual deve ser ignorado ligando o nó anterior diretamente ao próximo nó.

Para fazer isso, crie outra declaração `if` que verifique se previous não é estritamente `null`. Se for esse o caso, defina `previous.next` como `current.next`.

# --hints--

Você deve criar outra declaração `if` com a condição `previous !== null`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*/);
```

Você deve definir `previous.next` como `current.next`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /if\s*\(previous\s*!==\s*null\s*\)\s*\{\s*previous\.next\s*=\s*current\.next;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    while (current !== null && current.element !== element) {
        previous = current;
        current = current.next; 
    }

    if (current === null) {
        return;
    }
--fcc-editable-region--
    
--fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
