---
id: 69a014bd9b1730e94136279f
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Se a lista não estiver vazia, você precisa buscar o final dela, adicionar o novo elemento lá e fazer com que ele aponte para `null`. Isso significa que você deve fazer um loop pela lista a partir do `head`.

Para começar isso, abra uma cláusula `else`. Dentro dela, use `let` para criar uma variável `current` e definir seu valor como `list.head`.

# --hints--

Você deve ter uma cláusula `else`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.include(strippedAdd, 'else{');
```

Você deve criar uma variável `current` definida como `list.head` dentro da cláusula `else`.

```js
const explorer = await __helpers.Explorer(code);
const { add } = explorer.allFunctions;
const addSource = add?.toString() ?? '';
const withoutComments = __helpers.removeJSComments(addSource);
const strippedAdd = __helpers.removeWhiteSpace(withoutComments);
assert.match(strippedAdd, /else\{(?:let|var)current=list\.head/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList));
add(myList, 42);
console.log(myList);
console.log(isEmpty(myList));
```
