---
id: 69a012c6224fac85832247dd
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Percebe os resultados `true`, `{ head: null, length: 0 }` e `true` no console? Isso acontece porque você não atribuiu nada à lista `myList` e não incrementou a propriedade `length`.

Vá em frente e atribua `node` a `list.head` e, logo abaixo, incremente a propriedade `length` de `list`.

# --hints--

Você deve atribuir `node` a `list.head`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.head\s*=\s*node;?/);
```

Você deve incrementar a propriedade `length` de `list`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /list\.length\+\+;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    --fcc-editable-region--

    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
