---
id: 69a03ddce82b0626f329b001
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Agora, crie um loop `while` que verifique se `current` não é estritamente `null` e `current.element` não é estritamente igual ao elemento.

Isso significa que o loop deve continuar percorrendo a lista até o final dela ou até encontrar o elemento a ser removido.

# --hints--

Você deve criar um loop `while` com a condição `current !== null && current.element !== element`.

```js
assert.match(__helpers.removeJSComments(remove.toString()), /while\s+\(current\s*!==\s*null\s*&&\s*current\.element\s*!==\s*element\)\s*\{\s*/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        current.next = node;
    }

    list.length++;
}

function remove(list, element) {
    let previous = null;
    let current = list.head;

    --fcc-editable-region--
        
    --fcc-editable-region--
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
add(myList, 43);
add(myList, 44);
console.log(myList)
console.log(isEmpty(myList))
```
