---
id: 69a014bd9b1730e94136279f
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Se a lista não estiver vazia, você precisa buscar o final dela, adicionar o novo elemento lá e fazer com que ele aponte para `null`. Isso significa que você deve fazer um loop pela lista a partir do `head`.

Para começar isso, abra uma cláusula `else`. Dentro dela, use `let` para criar uma variável `current` e definir seu valor como `list.head`.

# --hints--

Você deve ter uma cláusula `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{/);
```

Você deve criar uma variável `current` definida como `list.head` dentro da cláusula `else`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /else\s*{\s*let|var\s+current\s*=\s*list\.head;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    --fcc-editable-region--
    } 
    --fcc-editable-region--

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
