---
id: 69a0102e761a245bb1cbd8e7
title: Step 5
challengeType: 1
dashedName: step-5
---

# --description--

Neste ponto, você deve testar a lista ligada para entender melhor.

Crie uma variável `myList` e defina-a como `initList()`. Logo abaixo, registre no console uma chamada de `isEmpty()` usando `myList` como parâmetro. Depois disso, chame `add` com `myList` e um número de sua escolha, então registre `myList` no console. Logo abaixo, registre outra chamada de `isEmpty()` no console usando `myList` como parâmetro novamente.

# --hints--

Você deve criar uma variável `myList` definida como `initList()`.

```js
assert.deepEqual(myList, initList());
```

Você deve registrar no console uma chamada para `isEmpty(myList)`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(isEmpty\(myList\)\);?/)
```

Você deve chamar `add` com `myList` e um número de sua escolha como parâmetros.

```js
assert.match(__helpers.removeJSComments(code), /add\s*\(\s*myList\s*,\s*\d+\s*\)/);
```

Você deve registrar sua variável `myList` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?/)
```

Você deve registrar `isEmpty(myList)` no console novamente.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(myList\);?\s*console\.log\(isEmpty\(myList\)\);?/)
```


# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };
}

--fcc-editable-region--

--fcc-editable-region--
```
