---
id: 69a01864d343dc2f2484d131
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Depois de encontrar o último nó, você precisa anexar o novo nó a ele. Para isso, defina `current.next` como `node` dentro do loop while, mas ainda dentro da cláusula `else`.

# --hints--

Você deve definir `current.next` como `node` dentro do seu loop `while`.

```js
assert.match(__helpers.removeJSComments(add.toString()), /current\.next\s*=\s*node;?/);
```

# --seed--

## --seed-contents--

```js
function initList() {
    return {
        head: null,
        length: 0
    };
}

function isEmpty(list) {
  return list.length === 0;
}

function add(list, element) {
    const node = {
        element: element,
        next: null
    };

    if (isEmpty(list)) {
        list.head = node;
    } else {
        let current = list.head;
        while (current.next !== null) {
            current = current.next;
        }
        --fcc-editable-region--

        --fcc-editable-region--
    }

    list.length++;
}

const myList = initList();
console.log(isEmpty(myList))
add(myList, 42);
console.log(myList)
console.log(isEmpty(myList))
```
