---
id: 69b279bae0f20eda63163cf8
title: Hatua 14
challengeType: 1
dashedName: step-14
---

# --description--

Kauli `if` na `else` ulizozitengeneza katika hatua zilizopita zitaweka vipengele katika safu ya data iliyopangwa.

Kila kipengele kilichopewa thamani katika safu ya data iliyopangwa kinachukua kielezo katika orodha. Kwa hivyo lazima uhamie kwenye kielezo kinachofuata katika safu ya data iliyopangwa baada ya kila mgawo wa thamani.

Chini ya kipande cha `if`/`else`, lakini bado ndani ya mzunguko wa `while`, ongeza `sortedIndex` kwa `1`.

# --hints--

Unapaswa kuongeza thamani ya sasa ya `sortedIndex` kwa `1` baada ya kifungu chako cha `else`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
