---
id: 69b2713ff6c21fed2d3ad88b
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Mpangilio wa merge ni algoriti ya kupanga inayotumia kanuni ya kugawanya-na-kushinda ili kupanga makusanyo ya data. Hiyo ni kusema, inagawa mkusanyo katika sehemu ndogo ndogo, na kushinda sehemu hizo ndogo kwa kuzipanga kwa uhuru, kisha kuunganisha sehemu zilizopangwa.

Anza kwa kutangaza kitendakazi kinachoitwa `mergeSort` kinachopokea kigezo `array`.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `mergeSort` chenye kigezo kimoja `array`.

```js
assert.isFunction(mergeSort);
const params = __helpers.getFunctionParams(mergeSort.toString());
assert.lengthOf(params, 1);
assert.deepEqual(params[0].name, "array");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
