---
id: 69b279bae0f20eda63163cf6
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Katika hatua iliyopita, uliweka thamani ya kipengee katika `leftPart` kwenye orodha ya `array` kwa sababu kilikuwa kidogo. Lakini hili halitakuwa daima hivyo. Katika baadhi ya kesi za kulinganisha, kipengee kilichopo upande wa kulia kinaweza kuwa kidogo.

Tengeneza kifungu cha `else` kitakachotekelezwa wakati kipengee katika `leftPart` si kidogo kuliko kipengee katika `rightPart`.

Ndani ya blokki ya `else`, weka thamani ya `rightPart[rightArrayIndex]` kwa `array[sortedIndex]`.

# --hints--

Unapaswa kutengeneza blokki ya `else` inayojumuisha `array[sortedIndex] = rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
