---
id: 69b279bae0f20eda63163cf9
title: Hatua 15
challengeType: 1
dashedName: step-15
---

# --description--

Mzunguko wa `while` uliouunda unalinganisha kipengele kimoja kutoka `leftPart` na kingine kutoka `rightPart`, kisha unaongeza kipengele kidogo zaidi kwenye orodha kuu ya `array`.

Utaendelea na operesheni hii hadi pasipokuwepo na vipengele vya kulinganisha. Lakini `leftPart` bado inaweza kuwa na vipengele wakati `rightPart` haina, na kinyume chake.

Unda mzunguko mwingine wa `while` unaoendesha wakati `leftArrayIndex` iko chini ya urefu wa `leftPart`. Katika hatua zinazofuata, utaitumia kunakili vipengele vilivyobaki katika `leftPart` kwenda kwenye orodha ya `array`.

# --hints--

Unapaswa kuunda mzunguko wa `while` wenye sharti la `leftArrayIndex < leftPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
