---
id: 69b279bae0f20eda63163cf2
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Ifuatayo, unahitaji kuunda mzunguko unaoendelea mradi tu vipengele bado vipo katika `leftPart` na `rightPart`.

Kwa ajili hiyo, unda mzunguko wa `while` wenye masharti mawili: moja linalokagua kama `leftArrayIndex` iko chini ya urefu wa `leftPart` na sharti lingine linalokagua kama `rightArrayIndex` iko chini ya urefu wa `rightPart`.

# --hints--

Unapaswa kuunda mzunguko wa `while` unaokagua kama `leftArrayIndex` iko chini ya urefu wa `leftPart` na kama `rightArrayIndex` iko chini ya urefu wa `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
