---
id: 69b2762b17d23fbee40a6c00
title: Step 2
challengeType: 1
dashedName: step-2
---

# --description--

Algoriti ya kupanga kwa kutumia merge sort hufanya vitendo vitatu:

- Gawanya mfululizo usiopangwa wa vitu katika sehemu ndogo ndogo
- Panga vitu katika sehemu ndogo ndogo
- Unganisha sehemu ndogo ndogo zilizopangwa

Haya yote hufanyika kwa mfululizo wa kurudia hadi sehemu ndogo ndogo ziumanishwe kuwa mfululizo mzima uliopangwa. Anza kwa kugawanya mfululizo.

Kwanza, ndani ya kitendakazi cha `mergeSort`, tangaza kigezo `middlePoint` kwa kutumia `const` na weka thamani `Math.floor(array.length / 2)` ili uweze kugawanya mfululizo nusu kabla ya kurudia.

# --hints--

Unapaswa kutangaza kigezo cha `middlePoint` na kuweka `Math.floor(array.length / 2)` kuwa thamani yake.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
