---
id: 69b279bae0f20eda63163cf4
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Wakati hali ya `if` inapotathminiwa kuwa `true`, inamaanisha kuwa kipengele kilicho katika orodha ya `leftPart` ni kidogo kuliko kipengele kinacholinganishwa nacho katika orodha ya `rightPart`.

Katika hali hiyo, unaweza kuweka kipengele kilichopatikana kwenye `leftArrayIndex` ndani ya `leftPart` katika safu ya data iliyopangwa.

Ndani ya blokki ya `if`, weka thamani ya `leftPart[leftArrayIndex]` kwa `array[sortedIndex]`.

# --hints--

Unapaswa kuweka thamani ya `leftPart[leftArrayIndex]` kwa `array[sortedIndex]` ndani ya `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
