---
id: 69b279bae0f20eda63163cf7
title: Hatua 13
challengeType: 1
dashedName: step-13
---

# --description--

Bado ukiwa ndani ya blokki ya `else`, ongeza `rightArrayIndex` kwa `1`.

# --hints--

Unapaswa kuongeza `rightArrayIndex` baada ya kuingiza kipengele katika safu ya data iliyopangwa.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
