---
id: 69b278519ef46b2fd8bf1546
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Tayari una thamani ya `middlePoint`. Tumia hiyo kugawanya safu ya data katika nusu mbili na weka sehemu ya kushoto tayari kwa kupanga.

Tumia sintaksia ya slice kutoa nusu ya kushoto ya `array` na weka thamani yake katika kigezo kinachoitwa `leftPart`.

# --hints--

Unapaswa kutangaza kigezo cha `leftPart` ndani ya kitendakazi cha `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

Unapaswa kuweka thamani ya `array.slice(0, middlePoint)` kwa `leftPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
