---
id: 69b2762b17d23fbee40a6c00
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Algoriti ya kupanga kwa kutumia merge sort hufanya vitendo vitatu:

- Gawanya mfululizo usiopangwa wa vitu katika sehemu ndogo ndogo
- Panga vitu katika sehemu ndogo ndogo
- Unganisha sehemu ndogo ndogo zilizopangwa

Haya yote hufanyika kwa mfululizo wa kurudia hadi sehemu ndogo ndogo ziumanishwe kuwa mfululizo mzima uliopangwa. Anza kwa kugawanya mfululizo.

Kwanza, ndani ya kitendakazi cha `mergeSort`, tangaza kigezo `middlePoint` kwa kutumia `const` na weka thamani `Math.floor(array.length / 2)` ili uweze kugawanya mfululizo nusu kabla ya kurudia.

# --hints--

Unapaswa kutangaza kigezo cha `middlePoint` na kuweka `Math.floor(array.length / 2)` kuwa thamani yake.

```js
const explorer = await __helpers.Explorer(code);

const { middlePoint } = explorer.allFunctions.mergeSort.variables;

assert.exists(
  middlePoint,
  'middlePoint should be declared inside mergeSort'
);

assert.isTrue(
  middlePoint.matches('const middlePoint = Math.floor(array.length / 2);'),
  'middlePoint should be declared with const and assigned Math.floor(array.length / 2)'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
