---
id: 69b279bae0f20eda63163cfe
title: Hatua 20
challengeType: 1
dashedName: step-20
---

# --description--

Kabla ya kujaribu kitendakazi cha `mergeSort()`, unahitaji kuunda kesi ya msingi inayosimamisha utekelezaji wa kitendakazi wakati urefu wa `array` ni chini au sawa na `1`.

Kesi hii ya msingi itasimamisha wito la mzunguko. Bila yake, operesheni ya kupanga kwa kutumia merge sort ingekuwa inaendelea hata wakati orodha imeshapangwa au haina vipengele ndani yake.

Mara moja baada ya tamko la kitendakazi, unda tamko la `if` linalokagua kama `array.length` ni chini au sawa na 1. Ndani ya mwili wa tamko la `if`, ongeza tamko la `return` kusimamisha utekelezaji wa kitendakazi.

# --hints--

Unapaswa kuunda tamko la `if` lenye sharti la `array.length <= 1`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

Unapaswa kuwa na tamko la `return` ndani ya mwili wa tamko la `if`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }
  
  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
