---
id: 69b2713ff6c21fed2d3ad88b
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Merge sort ni algoriti ya kupanga inayotumia kanuni ya kugawanya-na-kushinda kupanga makusanyo ya data. Hiyo ni, inagawa mkusanyo katika sehemu ndogo ndogo, na kushinda sehemu hizo kwa kuzipanga kwa uhuru, kisha kuunganisha sehemu zilizopangwa.

Anza kwa kutangaza kitendakazi kinachoitwa `mergeSort` kinachopokea kigezo `array`.

# --hints--

Unapaswa kutangaza kitendakazi kinachoitwa `mergeSort` chenye kigezo kimoja `array`.

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
