---
id: 69b279bae0f20eda63163cf5
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Baada ya kuweka thamani ya kipengele katika `leftPart` kwenye safu ya data iliyopangwa, ongeza `leftArrayIndex` kwa `1`.

# --hints--

Unapaswa kuongeza `1` kwa thamani ya sasa ya `leftArrayIndex` baada ya kuweka kipengele katika safu ya data iliyopangwa.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
