---
id: 69b279bae0f20eda63163cf1
title: Hatua 7
challengeType: 1
dashedName: step-7
---

# --description--

Sasa ni wakati wa kupanga na kuunganisha orodha (`leftPart` na `rightPart`) katika `array` ya awali.

Unaweza kufanya hivi kwa kulinganisha vipengele kwenye orodha zote mbili, na kuunganisha kipengele kidogo zaidi kwenye orodha kuu. Utafanya kulinganisha hii kwa kielezo zote katika `leftPart` na `rightPart`.

Tengeneza vigezo vitatu: `leftArrayIndex`, `rightArrayIndex`, na `sortedIndex`, na weka thamani zao kuwa `0`. Vigezo hivi vitakusaidia kufuatilia kila kielezo wakati wa mchakato wa kupanga.

# --hints--

Unapaswa kutangaza kigezo cha `leftArrayIndex` kilichoanzishwa kwa `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

Unapaswa kutangaza kigezo cha `rightArrayIndex` kilichoanzishwa kwa `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

Unapaswa kutangaza kigezo cha `sortedIndex` kilichoanzishwa kwa `0`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
