---
id: 69b279bae0f20eda63163cfd
title: Hatua 19
challengeType: 1
dashedName: step-19
---

# --description--

Mwishowe, unapaswa kunakili thamani zozote zilizobaki kutoka nusu ya kulia kwenda `array`.

Ili kufanya hivyo, ongeza mzunguko wa `while` unaoendelea kazi mradi `rightArrayIndex` ni ndogo kuliko `rightPart.length`, weka thamani ya `rightPart[rightArrayIndex]` ndani ya `array[sortedIndex]`, na ongeza thamani kwa `rightArrayIndex` na `sortedIndex`.

# --hints--

Unapaswa kuunda mzunguko wa `while` wenye sharti `rightArrayIndex < rightPart.length`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

Mwili wa mzunguko unapaswa kuhamisha `rightPart[rightArrayIndex]` kwenda `array[sortedIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

Unapaswa kuongeza thamani kwa `rightArrayIndex` na `sortedIndex` baada ya mgawo wa thamani.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }
  
  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
