---
id: 69b278d82f0e514f32acb724
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Baada ya kukata nusu ya kushoto, bado unahitaji upande wa kulia ili kukamilisha hatua ya kuunganisha.

Tumia sintaksia ya kukata (slice) kutoa nusu ya kulia ya `array` na weka thamani yake katika kigezo kinachoitwa `rightPart`.

# --hints--

Unapaswa kuwa na kigezo cha `rightPart` ndani ya kitendakazi cha `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

Unapaswa kutumia sintaksia ya kukata (slice) kutoa nusu ya kulia ya safu ya data na weka thamani yake katika kigezo cha `rightPart`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
