---
id: 69b279bae0f20eda63163cf0
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Chini kabisa ya mwili wa kitendakazi chako, fanya wito wa kitendakazi `mergeSort()` tena. Wakati huu, pita `rightPart` kama hoja katika wito wa kitendakazi.

# --hints--

Unapaswa kufanya wito wa `mergeSort()` ukiwa na hoja `rightPart` chini kabisa ya mwili wa kitendakazi chako.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
