---
id: 69b278d82f0e514f32acb724
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Baada ya kukata nusu ya kushoto, bado unahitaji upande wa kulia ili kukamilisha hatua ya kuunganisha.

Tumia sintaksia ya kukata (slice) kutoa nusu ya kulia ya `array` na weka thamani yake katika kigezo kinachoitwa `rightPart`.

# --hints--

Unapaswa kuwa na kigezo cha `rightPart` ndani ya kitendakazi cha `mergeSort`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
```

Unapaswa kutumia sintaksia ya kukata (slice) kutoa nusu ya kulia ya safu ya data na weka thamani yake katika kigezo cha `rightPart`.

```js
const explorer = await __helpers.Explorer(code);
const { rightPart } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightPart,
  'You should declare a rightPart variable inside mergeSort'
);
assert.isTrue(
  rightPart.value.matches('array.slice(middlePoint)'),
  'You should assign array.slice(middlePoint) to rightPart'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
