---
id: 69b279bae0f20eda63163cf3
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Ndani ya mzunguko, anza kulinganisha vipengele vya sasa kutoka kila nusu ili uweze kuunganisha thamani ndogo kwanza.

Ongeza kipengele cha `if` kinachokagua kama kipengele cha sasa katika `leftPart` ni kidogo kuliko kipengele cha sasa katika `rightPart`. 
# --hints--

Hali ya `if` inapaswa kukagua kama `leftPart[leftArrayIndex]` ni kidogo kuliko `rightPart[rightArrayIndex]`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
