---
id: 69b279bae0f20eda63163cef
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa baada ya kugawanya orodha ya `array` kuwa orodha mbili tofauti, utaendelea kugawanya kila orodha hadi kila kipengele kiwe katika orodha yake binafsi. Orodha yenye nambari moja tu huwa imepangwa tayari.

Ili kufanya hivyo, fanya wito wa `mergeSort` kwa mfululizo ndani ya kitendakazi chako na pitisha `leftPart` kama hoja kwa wito huo.

# --hints--

Unapaswa kuita `mergeSort(leftPart)` ndani ya kitendakazi hicho hicho cha `mergeSort`.

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
