---
id: 7a97b79c096ecb000d03bd36
title: Build an Nth Fibonacci Number Calculator
challengeType: 26
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**目标：** 完成以下用户故事并通过所有测试以完成实验。

**User Stories:**

1. 你应该创建一个名为 `fibonacci` 的函数。
2. 你应该在 `sequence` 函数内定义一个名为 `fibonacci` 的数组，并用 `[0, 1]` 的值初始化它。
3. `fibonacci` 函数应接受一个参数，一个正整数 `n`。
4. 调用 `fibonacci(n)` 应该使用动态规划方法计算并返回斐波那契序列中第 `n` 个数字，该序列中每个数字是前两个数字的和。
5. 斐波那契序列中位置为 `n` 的每个计算出的数字应保存在 `sequence` 数组的索引 `n - 1` 处。

# --hints--

你应该有一个名为 `fibonacci` 的函数。

```js
assert.isFunction(fibonacci);
```

你的 `fibonacci` 函数应接受一个参数。

```js
const params = __helpers.getFunctionParams(fibonacci.toString());
assert.equal(params.length, 1);
```

你应该在 `sequence` 函数内有一个名为 `fibonacci` 的数组，并用 `[0, 1]` 初始化。

```js
assert.match(__helpers.removeJSComments(fibonacci.toString()), /(const|let|var)\s+sequence\s*=\s*\[\s*0\s*,\s*1\s*\]/);
```

`fibonacci(0)` should return `0`.

```js
assert.equal(fibonacci(0), 0);
```

`fibonacci(1)` should return `1`.

```js
assert.equal(fibonacci(1), 1);
```

`fibonacci(2)` should return `1`.

```js
assert.equal(fibonacci(2), 1);
```

`fibonacci(3)` should return `2`.

```js
assert.equal(fibonacci(3), 2);
```

`fibonacci(5)` should return `5`.

```js
assert.equal(fibonacci(5), 5);
```

`fibonacci(10)` should return `55`.

```js
assert.equal(fibonacci(10), 55);
```

`fibonacci(15)` should return `610`.

```js
assert.equal(fibonacci(15), 610);
```

你的代码中不应使用递归。

```js
const funcStr = __helpers.removeJSComments(fibonacci.toString());
const bodyMatch = funcStr.match(/^[^{]*\{([\s\S]*)\}[^}]*$/) || funcStr.match(/=>\s*([\s\S]+)$/);
assert(bodyMatch && !bodyMatch[1].match(/\bfibonacci\s*\(/));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fibonacci(n) {
  const sequence = [0, 1];

  if (n < 2) {
    return sequence[n];
  }

  for (let i = 2; i <= n; i++) {
    sequence.push(sequence[i - 1] + sequence[i - 2]);
  }

  return sequence[n];
}
```
