---
id: 686383abac6c625a84458ea9
title: Schritt 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In diesem Workshop lernen Sie mehr über Tailwind CSS, indem Sie dessen Grid-Utility-Klassen verwenden, um eine Preis-Komponente zu erstellen.

Der grundlegende HTML-Boilerplate-Code wurde für Sie bereitgestellt. Beginnen Sie damit, ein `h1`-Element mit dem Text `Choose your listening plan` zu erstellen.

# --hints--

Sie sollten ein `h1`-Element erstellen.

```js
assert.exists(document.querySelector("h1"));
```

Ihr `h1`-Element sollte den Text `Choose your listening plan` enthalten.

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      --fcc-editable-region--
      
      --fcc-editable-region--
    </main>
  </body>
</html>
```
