---
id: 686386e76197458b7dfd52ef
title: Schritt 6
challengeType: 0
dashedName: step-6
---

# --description--

Jetzt erstellen Sie einen CTA-Button. Ein Call to Action (CTA) ist ein Button oder Link, der einen Benutzer dazu auffordert, eine bestimmte Aktion auszuführen, in diesem Fall mit dem Anhören zu beginnen.

Außerhalb des innersten `div` erstellen Sie ein Ankerelement mit einem `href` von `#` und dem Text `Start listening`.

Damit ist das HTML für die erste Preiskarte abgeschlossen.

# --hints--

Sie sollten ein Ankerelement erstellen.

```js
assert.exists(document.querySelector("a"))
```

Ihr Ankerelement sollte ein `href` mit dem Wert `#` haben.

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

Ihr Ankerelement sollte den Text `Start listening` enthalten.

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
