---
id: 686386e76197458b7dfd52ec
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Erstellen Sie unter dem `h2` ein `p`-Element mit dem Text `$0`. Dann erstellen Sie innerhalb dieses `p`-Elements und nach dem Text ein `span`-Element mit dem Text `/month`.

# --hints--

Sie sollten ein `p`-Element mit dem Text `$0` erstellen.

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

Sie sollten ein `span`-Element mit dem Text `/month` innerhalb des `p`-Elements erstellen.

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
