---
id: 686386e76197458b7dfd52eb
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Erstellen Sie ein `div`-Element unter dem `h1`. Innerhalb dieses `div` erstellen Sie ein weiteres `div`-Element und verschachteln darin ein weiteres `div`-Element.

Und innerhalb des innersten `div` erstellen Sie ein `h2`-Element mit dem Text `Listener`. Dies stellt den ersten Preisplan dar.

# --hints--

Sie sollten ein `div`-Element erstellen.

```js
assert.exists(document.querySelector("div"));
```

Sie sollten ein weiteres `div`-Element innerhalb des äußeren `div` erstellen.

```js
assert.exists(document.querySelector("div > div"));
```

Sie sollten ein weiteres `div`-Element innerhalb des inneren `div` erstellen.

```js
assert.exists(document.querySelector("div > div > div"));
```

Sie sollten ein `h2` mit dem Text `Listener` innerhalb des innersten `div`-Elements erstellen.

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
