---
id: 686386e76197458b7dfd52ee
title: Schritt 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstellen Sie innerhalb des `ul` zwei Listenelemente. Erstellen Sie im ersten Listenelement ein `span`-Element und fĂĽgen Sie die HTML-EntitĂ¤t `&#10003;` als Text hinzu. Direkt danach, aber noch innerhalb des Listenelements, fĂĽgen Sie den Text `span` ein, `Ad-supported streaming`.

Erstellen Sie im zweiten Listenelement ein `span`-Element und fĂĽgen Sie die EntitĂ¤t `&#10003;` als Text hinzu. Direkt danach, aber noch innerhalb des Listenelements, fĂĽgen Sie den Text `span` ein, `Curated playlists`.

# --hints--

Sie sollten zwei Listenelemente innerhalb Ihres `ul`-Elements haben.

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

Ihr erstes `li`-Element sollte ein `span`-Element mit der HTML-ZeichenentitĂ¤t `&#10003;` enthalten.

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "âś“")
```

Ihr erstes `li`-Element sollte den Text `Ad-supported streaming` direkt nach dem `span`-Element enthalten.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

Ihr zweites `li`-Element sollte ein `span`-Element mit der HTML-ZeichenentitĂ¤t `&#10003;` enthalten.

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "âś“")
```

Ihr zweites `li`-Element sollte den Text `Curated playlists` direkt nach dem `span`-Element enthalten.

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
