---
id: 6864d9999f8a4e3f0d569a3f
title: Schritt 8
challengeType: 0
dashedName: step-8
---

# --description--

Machen Sie als Nächstes den `h1`-Text unterscheidbar, indem Sie ihm eine `class` zuweisen, die auf `text-3xl` gesetzt ist, um ihn größer zu machen, `text-center`, um ihn zu zentrieren, `font-semibold`, um ihn fetter zu machen, `text-gray-900`, um ihn dunkler zu machen, und `md:text-5xl`, um ihn auf mittleren und größeren Bildschirmgrößen größer zu machen.

# --hints--

Ihr `h1`-Element sollte die Klasse `text-3xl` haben.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-3xl"))
```

Ihr `h1`-Element sollte die Klasse `text-center` haben.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-center"))
```

Ihr `h1`-Element sollte die Klasse `font-semibold` haben.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("font-semibold"))
```

Ihr `h1`-Element sollte die Klasse `text-gray-900` haben.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("text-gray-900"))
```

Ihr `h1`-Element sollte die Klasse `md:text-5xl` haben.

```js
const h1El = document.querySelector("h1")

assert.isTrue(h1El.classList.contains("md:text-5xl"))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body class="bg-gray-100">
    <main>
      --fcc-editable-region--
      <h1>Choose your listening plan</h1>
      --fcc-editable-region--
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          <a href="#">Start listening</a>
        </div>
        <div>
          <div>Most Popular</div>
          <div>
            <h2>Premium</h2>
            <p>$9.99<span>/month</span></p>
            <p>
              Enjoy the full music experience with unlimited access and
              downloads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-free listening</li>
              <li><span>&#10003;</span>Offline playback</li>
              <li><span>&#10003;</span>Unlimited skips</li>
            </ul>
          </div>
          <a href="#">Go Premium</a>
        </div>
        <div>
          <div>
            <h2>Family</h2>
            <p>$14.99<span>/month</span></p>
            <p>
              Enjoy all of the features with a plan for up to 6 family members.
            </p>
            <ul>
              <li>
                <span>&#10003;</span>All Premium features
              </li>
              <li>
                <span>&#10003;</span>Up to 6 accounts
              </li>
              <li>
                <span>&#10003;</span>Individual playlists &
                libraries
              </li>
              <li>
                <span>&#10003;</span>Family Mix playlists
              </li>
            </ul>
          </div>
          <a href="#">Start Family Plan</a>
        </div>
      </div>
    </main>
  </body>
</html>
```
