---
id: 66eaddd04a9e533fba689001
title: Створіть картку для допису
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

У цій лабораторній роботі ви попрактикуєтесь стилізувати фон та кордон, створивши картку для допису блогу.

**Мета:** виконайте історію користувача, щоб завершити лабораторну роботу. Усі тести мають бути успішними.

**Історія користувача:**

1. Ви повинні мати елемент `div` з класом `blog-post-card`, який містить всі елементи картки.
2. В межах елемента `.blog-post-card` має бути зображення з дійсним атрибутом `alt` і текстом, а також класом `post-img`. Для атрибута `src` можна використати `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg`.
3. В межах елемента `.blog-post-card` має бути елемент `div` з класом `post-content` і таким вмістом:
    - елемент `h2` з класом `post-title` та заголовком допису;
    - елемент `p` з класом `post-excerpt` та коротким описом допису;
    - елемент `a` з класом `read-more` та текстом `Read More`.
4. До елемента `.blog-post-card` застосуйте таку стилізацію:
    - білий фон;
    - заокруглені кути;
    - ширина на ваш вибір;
    - вирівнювання тексту на ваш вибір.
5. Елемент `.post-img` має бути стилізований так, щоб зображення заповнювало всю ширину картки та мало нижній кордон.
6. Елемент `.post-content` має бути стилізований так, щоб всередині був відступ.
7. Елементи `.post-title` та `.post-excerpt` повинні мати нестандартний колір тексту, а також поля з усіх сторін.
8. Елемент `.read-more` повинен бути стилізований як кнопка та мати:
    - нестандартний колір тексту;
    - колір фону;
    - поля з усіх сторін;
    - відступи з усіх сторін;
    - заокруглені кути;
    - властивість `display` зі значенням `inline-block`;
    - ефект при наведенні курсора, що змінює колір фону.

**Примітка:** не забудьте прив’язати таблицю стилів до HTML та застосувати CSS.

# --hints--

Ви повинні мати елемент `div` з класом `blog-post-card`.

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

Ви повинні мати елемент `img` з `class` зі значенням `post-img`. Переконайтесь, що зображення має атрибут `alt` з текстом та атрибут `src` зі значенням.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

Ви повинні мати елемент `div` з класом `post-content`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

Ви повинні мати елемент `h2` з класом `post-title`. Переконайтесь, що він містить заголовок допису.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

Ви повинні мати елемент `p` з класом `post-excerpt`. Переконайтесь, що він містить короткий опис допису.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

Ви повинні мати елемент `a` з `class` зі значенням `read-more`.

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

Елемент `.read-more` повинен мати текст `Read More`.

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent, "Read More");
```

Елемент `.blog-post-card` повинен мати властивість `border-radius`, значення якої не може бути `0` або від’ємним.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

Елемент `.blog-post-card` повинен мати білий фон.

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

Націльтесь на елемент `.blog-post-card` і встановіть його властивість `width`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

Націльтесь на елемент `.blog-post-card` і встановіть його властивість `text-align`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

Націльтесь на елемент `.post-content` і встановіть його властивість `padding`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

Елемент `.read-more` повинен мати ефект при наведенні курсора.

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

Націльтесь на елемент `.read-more` і встановіть його властивість `color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

Націльтесь на елемент `.read-more` і встановіть його властивість `background-color`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

Націльтесь на елемент `.read-more` і встановіть його властивість `margin`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

Націльтесь на елемент `.read-more` і встановіть його властивість `display`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

Націльтесь на елемент `.read-more` і встановіть його властивість `border-radius`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

Націльтесь на елемент `.read-more` і встановіть його властивість `padding`.

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

Елемент `.post-img` повинен заповнювати ширину картки.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

Елемент `.post-img` повинен мати властивість `border-bottom`.

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

Елементи `.post-title` та `.post-excerpt` повинні мати поля, а також нестандартні кольори тексту.

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
