---
id: 6140c7e645d8e905819f1dd4
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

Añade tu elemento `meta`, `charset` para el conjunto de caracteres correcto, tu elemento `title`, y un elemento `link` para el archivo `./styles.css`.

Establece el `title` a `Ferris Wheel`.

# --hints--

Tu código debe contener la declaración `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Debe incluir un espacio después de la declaración `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Debes definir el tipo de documento a `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Debes cerrar la declaración `DOCTYPE` con un `>` después del tipo.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Tu elemento `html` debe tener una etiqueta de apertura con un atributo `lang` con el valor `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

El elemento `html` debe tener un tag de cierre.

```js
assert.match(code, /<\/html\s*>/);
```

La declaración `DOCTYPE` debe estar al comienzo del documento HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Debe haber un tag de apertura para `head`.

```js
assert.match(code, /<head\s*>/i);
```

Debe haber un tag de cierre para `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Debe haber un tag de apertura para `body`.

```js
assert.match(code, /<body\s*>/i);
```

Debe haber un tag de cierre para `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

El elemento `body` debe ser el primero después del elemento `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Los elementos de la `head` y el `body` deben ser hermanos.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Tu elemento `head` debe ser el primer elemento dentro de tu elemento `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Tu elemento `head` debe estar dentro del elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Tu elemento `body` debe estar dentro del elemento `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Your code should have a `meta` element.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Your `meta` element should have a `charset` attribute with the value `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Tu elemento `meta` debe estar dentro del elemento `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

Debes tener un elemento `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
