---
id: 684aaf9ec670c68d20efd0d0
title: Construir um Gerenciador de Configuração de Usuário
challengeType: 27
dashedName: build-a-user-configuration-manager
---

# --description--

Neste laboratório, você construirá um Gerenciador de Configuração de Usuário que permite aos usuários gerenciar suas configurações, como tema, idioma e notificações. Você implementará funções para adicionar, atualizar, excluir e visualizar as configurações do usuário.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve definir uma função chamada `add_setting` com dois parâmetros que representam um dicionário de configurações e uma tupla contendo um par chave-valor

1. A função `add_setting` deve:
    
   - Converta a chave e o valor para letras minúsculas.
   - Se a configuração da chave existir, retorne `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - Se a configuração da chave não existir, adicione o par chave-valor ao dicionário de configurações fornecido e retorne `Setting '[key]' added with value '[value]' successfully!`.
   - As mensagens retornadas devem ter a chave e o valor em letras minúsculas.

1. Você deve definir uma função chamada `update_setting` com dois parâmetros que representam um dicionário de configurações e uma tupla contendo um par chave-valor.

1. A função `update_setting` deve:
       
   - Converta a chave e o valor para letras minúsculas.
   - Se a configuração da chave existir, atualize seu valor no dicionário de configurações fornecido e retorne: `Setting '[key]' updated to '[value]' successfully!`
   - Se a configuração da chave não existir, retorne `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - As mensagens retornadas devem ter a chave e o valor em letras minúsculas.


1. Você deve definir uma função chamada `delete_setting` com dois parâmetros que representam um dicionário de configurações e uma chave.

1. A função `delete_setting` deve:
    
   - Converta a chave passada para letras minúsculas.
   - Se a configuração da chave existir, remova o par chave-valor do dicionário de configurações fornecido e retorne `Setting '[key]' deleted successfully!`
   - Se a configuração da chave não existir, retorne `Setting not found!`
   - As mensagens retornadas devem ter a chave em letras minúsculas.

1. Você deve definir uma função chamada `view_settings` com um parâmetro que representa um dicionário de configurações.

1. A função `view_settings` deve:
   
   - Retorne `No settings available.` se o dicionário de configurações fornecido estiver vazio. 
   - Se o dicionário contiver alguma configuração, retorne uma string exibindo as configurações. A string deve começar com `Current User Settings:` seguida dos pares chave-valor, cada um em uma nova linha e com a chave capitalizada. Por exemplo, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` deve retornar:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Para testar o código, você deve criar um dicionário chamado `test_settings` para armazenar algumas preferências de configuração do usuário.

# --hints--

Você deve criar um dicionário chamado `test_settings` e adicionar alguns valores a ele.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Você deve definir uma função chamada `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

A função `add_setting` deve ter dois parâmetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` deve converter a chave para letras minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` deve converter o valor para letras minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` deve retornar a mensagem de erro `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` deve adicionar um novo par chave-valor e retornar a mensagem de sucesso `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` deve adicionar corretamente o par chave-valor fornecido ao dicionário.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Você deve definir uma função chamada `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

A função `update_setting` deve ter dois parâmetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

A função `update_setting` deve converter a chave para letras minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

A função `update_setting` deve converter o valor para letras minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` deve atualizar uma chave existente e retornar a mensagem de sucesso `Configuração 'theme' atualizada para 'dark' com sucesso!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` deve retornar a mensagem de erro `Configuração 'volume' não existe!` Não é possível atualizar uma configuração inexistente.` quando a chave não existe.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` deve atualizar corretamente o par chave-valor fornecido no dicionário.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Você deve definir uma função chamada `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

A função `delete_setting` deve ter dois parâmetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` deve converter a chave para letras minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` deve remover a chave existente e retornar a mensagem de sucesso `Configuração 'theme' excluída com sucesso!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` deve retornar a mensagem de erro `Configuração não encontrada!` quando a chave não existir.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` deve remover corretamente a chave fornecida do dicionário.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Você deve definir uma função chamada `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

A função `view_settings` deve ter um parâmetro.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` deve retornar a mensagem `Nenhuma configuração disponível.` se o dicionário fornecido estiver vazio.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` deve retornar configurações formatadas para dicionário não vazio.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` deve capitalizar a primeira letra de cada nome de configuração.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` deve exibir os resultados corretos e terminar com um caractere de nova linha.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: English\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
