---
id: 587d825a367417b2b2512c89
title: реалізувати швидке сортування
challengeType: 1
forumTopicId: 301615
dashedName: implement-quick-sort
---

# --description--

Тепер ми перейдемо до проміжного алгоритму сортування: швидкого сортування. Швидке сортування — це ефективний рекурсивний підхід «розділяй і володарюй» для сортування масиву. У цьому методі в оригінальному масиві вибирається опорне значення. Потім масив розбивається на два підмасиви: зі значеннями меншими та більшими за опорне. Далі ми об’єднуємо результат рекурсивного виклику алгоритму швидкого сортування для обох підмасивів. Це триває, доки не буде досягнуто базового випадку — порожнього або однеелементного масиву, який ми повертаємо. Розгортання рекурсивних викликів повертає нам відсортований масив.

Швидке сортування — дуже ефективний метод сортування, який у середньому має продуктивність *O(nlog(n))*. Також його досить просто реалізувати. Ці характеристики роблять його популярним і корисним методом сортування.

**Інструкції:** Напишіть функцію `quickSort`, яка приймає масив цілих чисел як вхідні дані і повертає масив цих чисел, відсортований за зростанням. Хоча вибір опорного значення важливий, для наших цілей підійде будь-яке опорне. Для простоти можна використовувати перший або останній елемент.

# --hints--

`quickSort` має бути функцією.

```js
assert(typeof quickSort == 'function');
```

`quickSort` має повертати відсортований масив (від найменшого до найбільшого).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    quickSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`quickSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` має повертати масив, який не змінений, крім порядку елементів.

```js
assert.sameMembers(
  quickSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`quickSort` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    quickSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function quickSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function quickSort(array) {
  if (array.length === 0) {
    return [];
  } else {
    const pivotValue = array[0];

    // Sort elements into three piles
    let lesser = [];
    let equal = [];
    let greater = [];
    for (let e of array) {
      if (e < pivotValue) {
        lesser.push(e);
      } else if (e > pivotValue) {
        greater.push(e);
      } else {
        equal.push(e);
      }
    }

    return [...quickSort(lesser), ...equal, ...quickSort(greater)];
  }
}
```
