---
id: 587d825c367417b2b2512c8f
title: Реалізувати сортування злиттям
challengeType: 1
forumTopicId: 301614
dashedName: implement-merge-sort
---

# --description--

Ще одним загальним алгоритмом сортування середнього рівня є сортування злиттям. Як і швидке сортування, сортування злиттям також використовує метод розділяй і володарюй, рекурсивно сортує масив. Воно використовує той факт, що відносно легко відсортувати два масиви, якщо кожен із них вже відсортований. Але спочатку у нас є лише один масив на вхід, тож як отримати два відсортовані масиви з нього? Ми можемо рекурсивно розділяти початковий масив навпіл, доки не дійдемо до базового випадку — масиву з одним елементом. Масив з одним елементом природно відсортований, тож потім можна починати об’єднувати. Це об’єднання розгорне рекурсивні виклики, які розбивали початковий масив, і врешті-решт утворить фінальний відсортований масив усіх елементів. Кроки сортування злиттям такі:

**1)** Рекурсивно розбити вхідний масив навпіл, доки не утвориться підмасив з одним елементом.

**2)** Об’єднати кожен відсортований підмасив, щоб отримати фінальний відсортований масив.

Сортування злиттям — це ефективний метод сортування з часовою складністю *O(nlog(n))*. Цей алгоритм популярний, бо він продуктивний і відносно простий для реалізації.

До речі, це буде останній алгоритм сортування, який ми розглянемо тут. Однак пізніше, у розділі про структури даних у вигляді дерев, ми опишемо сортування купою — ще один ефективний метод сортування, який потребує двійкової купи для реалізації.

**Інструкції:** Напишіть функцію `mergeSort`, яка приймає масив цілих чисел на вхід і повертає масив цих чисел, відсортований за зростанням. Добрий спосіб реалізувати це — написати одну функцію, наприклад `merge`, яка відповідає за об’єднання двох відсортованих масивів, і іншу функцію, наприклад `mergeSort`, яка відповідає за рекурсію, що утворює масиви з одним елементом для подачі в об’єднання. Успіхів!

# --hints--

`mergeSort` має бути функцією.

```js
assert.isFunction(mergeSort);
```

`mergeSort` має повертати відсортований масив (від найменшого до найбільшого).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}
assert.isTrue(
  isSorted(
    mergeSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`mergeSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` має повертати масив, який не змінений, окрім порядку елементів.

```js
assert.sameMembers(
  mergeSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`mergeSort` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    mergeSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function mergeSort(array) {
  if (array.length === 1) {
    return array;
  } else {
    const splitIndex = Math.floor(array.length / 2);
    return merge(
      mergeSort(array.slice(0, splitIndex)),
      mergeSort(array.slice(splitIndex))
    );
  }

  // Merge two sorted arrays
  function merge(array1, array2) {
    let merged = [];
    while (array1.length && array2.length) {
      if (array1[0] < array2[0]) {
        merged.push(array1.shift());
      } else if (array1[0] > array2[0]) {
        merged.push(array2.shift());
      } else {
        merged.push(array1.shift(), array2.shift());
      }
    }

    // After looping ends, one array is empty, and other array contains only
    // values greater than all values in `merged`
    return [...merged, ...array1, ...array2];
  }
}

mergeSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]);
```
