---
id: 587d8259367417b2b2512c85
title: реалізувати сортування вибором
challengeType: 1
forumTopicId: 301616
dashedName: implement-selection-sort
---

# --description--

Тут ми реалізуємо сортування вибором. Сортування вибором працює шляхом вибору мінімального значення в списку та обміну його з першим значенням у списку. Потім воно починає з другої позиції, вибирає найменше значення в залишку списку і обмінює його з другим елементом. Процес продовжується, ітеруючись по списку та обмінюючи елементи, доки не досягне кінця списку. Тепер список відсортований. Сортування вибором має квадратичну часову складність у всіх випадках.

**Інструкції**: Напишіть функцію `selectionSort`, яка приймає масив цілих чисел як вхідні дані і повертає масив цих чисел, відсортований за зростанням.

# --hints--

`selectionSort` має бути функцією.

```js
assert.isFunction(selectionSort);
```

`selectionSort` має повертати відсортований масив (від найменшого до найбільшого).

```js
function isSorted(a){
  for(let i = 0; i < a.length - 1; i++)
    if(a[i] > a[i + 1])
      return false;
  return true;
}

assert.isTrue(
  isSorted(
    selectionSort([
      1,
      4,
      2,
      8,
      345,
      123,
      43,
      32,
      5643,
      63,
      123,
      43,
      2,
      55,
      1,
      234,
      92
    ])
  )
);
```

`selectionSort([1,4,2,8,345,123,43,32,5643,63,123,43,2,55,1,234,92])` має повертати масив, який не змінений, окрім порядку елементів.

```js
assert.sameMembers(
  selectionSort([
    1,
    4,
    2,
    8,
    345,
    123,
    43,
    32,
    5643,
    63,
    123,
    43,
    2,
    55,
    1,
    234,
    92
  ]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

`selectionSort` не має використовувати вбудований метод `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js
function selectionSort(array) {
  // Only change code below this line
  return array;
  // Only change code above this line
}
```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
