---
id: 61abc7ebf3029b56226de5b6
title: Реалізувати бінарний пошук
challengeType: 1
forumTopicId: 487618
dashedName: implement-binary-search
---

# --description--

Бінарний пошук — це алгоритм з ефективністю **O(log(n))** для пошуку елемента у відсортованому масиві. Він працює за такими кроками:

1. Знайти середній `value` відсортованого масиву. Якщо `value == target`, повернути `true` (значення знайдено, пошук завершено).
1. Якщо середній `value < target`, шукати у правій половині масиву при наступному порівнянні.
1. Якщо середній `value > target`, шукати у лівій половині масиву при наступному порівнянні.
1. Якщо після пошуку всього масиву значення відсутнє, повернути `false` (масив було переглянуто, значення в ньому немає).

Як бачите, ви послідовно ділите масив навпіл, що дає ефективність log(n). Для цього завдання ми хочемо, щоб ви показали свій шлях — як ви дійшли до цільового значення... шлях, яким ви йшли!

# --instructions--

Напишіть функцію `binarySearch`, яка реалізує алгоритм бінарного пошуку у масиві, повертаючи шлях, яким ви йшли (кожне порівняння з середнім значенням), щоб знайти ціль у масиві.

Функція приймає відсортований масив цілих чисел і цільове значення як вхідні дані. Вона повертає масив, що містить (у порядку зростання) середнє значення, яке ви знаходили на кожному кроці поділу початкового масиву, доки не знайшли цільове значення. Цільове значення має бути останнім елементом повернутого масиву. Якщо значення не знайдено, поверніть рядок `Value Not Found`.

Наприклад, `binarySearch([1,2,3,4,5,6,7], 5)` поверне `[4,6,5]`.

Для цього завдання при поділі ви ОБОВ’ЯЗКОВО маєте використовувати `Math.floor()` при діленні: `Math.floor(x/2)`. Це забезпечить послідовний, тестований шлях.

**Примітка:** Для тестів буде використано наступний масив:

```js
const testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
```

# --hints--

`binarySearch` має бути функцією.

```js
assert.isFunction(binarySearch);
```

`binarySearch(testArray, 0)` має повертати `[13, 5, 2, 0]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 0), [13, 5, 2, 0]);
```

`binarySearch(testArray, 1)` має повертати `[13, 5, 2, 0, 1]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 1), [13, 5, 2, 0, 1]);
```


`binarySearch(testArray, 2)` має повертати `[13, 5, 2]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 2), [13, 5, 2]);
```

`binarySearch(testArray, 6)` має повертати рядок `Value Not Found`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.strictEqual(binarySearch(_testArray, 6), 'Value Not Found');
```

`binarySearch(testArray, 11)` має повертати `[13, 5, 10, 11]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 11), [13, 5, 10, 11])
```

`binarySearch(testArray, 13)` має повертати `[13]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 13), [13]);
```

`binarySearch(testArray, 70)` має повертати `[13, 19, 22, 49, 70]`.

```js
const _testArray = [
  0, 1, 2, 3, 4, 5, 8, 9, 10, 11, 12, 13, 14, 15, 16, 17, 18, 19, 20, 21, 22,
  23, 49, 70
];
assert.deepEqual(binarySearch(_testArray, 70), [13, 19, 22, 49, 70]);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let arrayPath = [];
  return arrayPath;
}
```



# --solutions--

```js
let binarySearch = (searchList, value) => {
  let arrayPath = [];

  // set initial L - M - R
  let left = 0;
  let right = searchList.length - 1;
  let middle = Math.floor(right / 2);

  // if first comparison finds value
  if (searchList[middle] == value) {
    arrayPath.push(searchList[middle]);
    return arrayPath;
  }

  while (searchList[middle] !== value) {
    // add to output array
    arrayPath.push(searchList[middle]);

    // not found
    if (right < left) {
      return 'Value Not Found';
    }
    // value is in left or right portion of array
    // update L - M - R
    if (searchList[middle] > value) {
      right = middle - 1;
      middle = left + Math.floor((right - left) / 2);
    } else {
      left = middle + 1;
      middle = left + Math.floor((right - left) / 2);
    }

    // if found update output array and exit
    if (searchList[middle] == value) {
      arrayPath.push(searchList[middle]);

      break;
    }
  }
  return arrayPath;
};
```
