---
id: 6976db7ecfa770bf21307b20
title: 构建播放列表混音引擎
challengeType: 26
dashedName: build-a-playlist-remix-engine
---

# --description--

在本实验中，你将构建一个程序，从多个听众提交的播放列表中创建一个混音播放列表。

每个听众都会提供他们想听的歌曲列表。有些歌曲可能出现多次，有些艺术家可能出现次数过多。你的任务是逐步处理这些播放列表：将它们合并成一个列表，为每首歌评分，删除重复歌曲，限制同一艺术家的出现次数，然后创建最终的播放顺序。

**目标：** 完成下面的用户故事并通过所有测试以完成实验。

**用户故事：**

1. 你应该创建一个名为 `flattenPlaylists` 的函数，该函数接受一个播放列表数组，其中每个播放列表是包含以下属性的对象数组：`trackId`、`artist`、`title`、`votes`、`bpm`。如果输入不是数组，`flattenPlaylists` 应该返回一个空数组。已为你提供了一个示例播放列表，你可以用它来测试你的函数。

2. `flattenPlaylists` 应该返回一个扁平的曲目对象数组，每个对象包含所有原始曲目属性以及一个 `source` 属性，该属性是一个数组，包含播放列表索引和曲目索引，表示曲目的来源位置。

3. 你应该创建一个名为 `scoreTracks` 的函数，该函数接受一个由 `flattenPlaylists` 返回的曲目对象数组（每个对象具有 `trackId`、`artist`、`title`、`votes`、`bpm` 和 `source` 属性），并返回一个新的曲目对象数组，每个对象添加了一个 `score` 属性，计算公式为：`votes * 10 - Math.abs(bpm - 120)`。

4. 你应该创建一个名为 `dedupeTracks` 的函数，该函数接受一个由 `scoreTracks` 返回的曲目对象数组，并返回一个新的数组，删除重复的 `trackId` 条目，只保留第一次出现的。

5. 你应该创建一个名为 `enforceArtistQuota` 的函数，该函数接受一个由 `dedupeTracks` 返回的曲目对象数组和一个表示每个艺术家最大允许出现次数的数字。该函数应返回一个新数组，确保没有艺术家出现次数超过给定数字，保留最早出现的。

6. 你应该创建一个名为 `buildSchedule` 的函数，该函数接受一个由 `enforceArtistQuota` 返回的曲目对象数组，并返回一个新的 `{ slot, trackId }` 对象数组，其中 `slot` 是一个从 1 开始的索引，表示每首曲目在播放顺序中的位置。

7. 你应该创建一个名为 `remixPlaylist` 的函数，该函数接受一个播放列表数组和每个艺术家允许的最大出现次数。该函数应通过依次调用 `{ slot, trackId }`、`flattenPlaylists`、`scoreTracks`、`dedupeTracks` 和 `enforceArtistQuota`，返回最终的播放调度数组，数组元素为 `buildSchedule` 对象。

# --hints--

你应该有一个名为 `flattenPlaylists` 的函数。

```js
assert.isFunction(flattenPlaylists);
```

当输入不是数组时，`flattenPlaylists` 应该返回一个空数组。

```js
assert.deepEqual(flattenPlaylists(null), []);
assert.deepEqual(flattenPlaylists(undefined), []);
assert.deepEqual(flattenPlaylists({}), []);
assert.deepEqual(flattenPlaylists("not an array"), []);
assert.deepEqual(flattenPlaylists(123), []);
```

`flattenPlaylists` 返回的每个曲目应包含一个 `source` 字段，该字段是一个数组，包含播放列表索引和曲目索引。

```js
const playlists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 110 }
  ],
  [{ trackId: "t3", artist: "C", title: "Song 3", votes: 5, bpm: 130 }]
];

const flat = flattenPlaylists(playlists);

assert.isArray(flat);
assert.lengthOf(flat, 3);

flat.forEach((t) => {
  assert.property(t, "source");
  assert.isArray(t.source);
  assert.lengthOf(t.source, 2);
  assert.isNumber(t.source[0]);
  assert.isNumber(t.source[1]);
});

assert.deepEqual(flat[0].source, [0, 0]);
assert.deepEqual(flat[1].source, [0, 1]);
assert.deepEqual(flat[2].source, [1, 0]);
```

你应该有一个名为 `scoreTracks` 的函数。

```js
assert.isFunction(scoreTracks)
```

`scoreTracks` 返回的每个曲目应包含一个数值类型的 `score` 字段。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120, source: [0, 0] }
];

const scored = scoreTracks(tracks);

assert.isArray(scored);
assert.lengthOf(scored, 1);

assert.property(scored[0], "score");
assert.isNumber(scored[0].score);
```

你应该使用目标 BPM 为 `score`，并用公式 `120` 计算 `votes * 10 - Math.abs(bpm - 120)`。

```js
const tracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 3, bpm: 120, source: [0, 0] }, // 30
  { trackId: "t2", artist: "B", title: "Song 2", votes: 1, bpm: 100, source: [0, 1] },  // -10
  { trackId: "t3", artist: "C", title: "Song 3", votes: 2, bpm: 135, source: [1, 0] }  // 5
];

const formulaScored = scoreTracks(tracks);

const s1 = formulaScored.find(t => t.trackId === "t1").score;
const s2 = formulaScored.find(t => t.trackId === "t2").score;
const s3 = formulaScored.find(t => t.trackId === "t3").score;

assert.equal(s1, 30);
assert.equal(s2, -10);
assert.equal(s3, 5);
```

你应该有一个名为 `dedupeTracks` 的函数。

```js
assert.isFunction(dedupeTracks)
```

当存在重复的 `trackId` 值时，`dedupeTracks` 应只保留曲目的第一次出现。

```js
const dupes = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "B", title: "Song 2", votes: 2, bpm: 119, source: [0, 1], score: 19 },
  { trackId: "t1", artist: "A", title: "Song 3", votes: 9, bpm: 140, source: [1, 0], score: 70 }
];

const deduped = dedupeTracks(dupes);

assert.isArray(deduped);
assert.lengthOf(deduped, 2);

assert.equal(deduped[0].trackId, "t1");
assert.equal(deduped[0].title, "Song 1");
assert.equal(deduped[1].trackId, "t2");
assert.equal(deduped[1].title, "Song 2");

```

你应该有一个名为 `enforceArtistQuota` 的函数。

```js
assert.isFunction(enforceArtistQuota)
```

`enforceArtistQuota` 应确保没有艺术家出现次数超过 `maxPerArtist`，通过删除多余曲目并保留最早出现的。

```js
const quotaTracks = [
  { trackId: "t1", artist: "A", title: "Song 1", votes: 1, bpm: 120, source: [0, 0], score: 10 },
  { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 121, source: [0, 1], score: 9 },
  { trackId: "t3", artist: "B", title: "Song 3", votes: 1, bpm: 118, source: [0, 2], score: 8 },
  { trackId: "t4", artist: "A", title: "Song 4", votes: 1, bpm: 110, source: [1, 0], score: 0 }
];

const limited = enforceArtistQuota(quotaTracks, 2);

let artistACount = 0;

for (const track of limited) {
  if (track.artist === "A") {
    artistACount++;
  }
}
assert.isAtMost(artistACount, 2);
assert.notInclude(limited.map(t => t.trackId), "t4");
```

你应该有一个名为 `buildSchedule` 的函数。

```js
assert.isFunction(buildSchedule);
```

`buildSchedule` 应返回一个对象数组，形状为 `{ slot, trackId }`，其中 `slot` 从 `1` 开始。

```js
const scheduleInput = [{ trackId: "t1" }, { trackId: "t2" }, { trackId: "t3" }];
const schedule = buildSchedule(scheduleInput);

assert.lengthOf(schedule, 3);

schedule.forEach((item, index) => {
  assert.hasAllKeys(item, ["slot", "trackId"]);
  assert.equal(item.slot, index + 1);
  assert.equal(item.trackId, scheduleInput[index].trackId);
});
```

你应该有一个名为 `remixPlaylist` 的函数。

```js
assert.isFunction(remixPlaylist)
```

`remixPlaylist` 应按顺序调用辅助函数以生成最终的播放调度。

```js
assert.lengthOf(remixPlaylist, 2);

const testPlaylists = [
  [
    { trackId: "t1", artist: "A", title: "Song 1", votes: 2, bpm: 120 },
    { trackId: "t1", artist: "A", title: "Duplicate", votes: 9, bpm: 140 }
  ],
  [
    { trackId: "t2", artist: "A", title: "Song 2", votes: 1, bpm: 110 },
    { trackId: "t3", artist: "B", title: "Song 3", votes: 4, bpm: 118 }
  ]
];

const expected = buildSchedule(enforceArtistQuota(dedupeTracks(scoreTracks(flattenPlaylists(testPlaylists))),1));

const actual = remixPlaylist(testPlaylists, 1);

assert.deepEqual(actual, expected);
```

# --seed--

## --seed-contents--

```js
const playlists = [
  [
    {
      trackId: "trk101",
      artist: "Velvet Comet",
      title: "Crimson Afterglow",
      votes: 5,
      bpm: 122
    },
    {
      trackId: "trk102",
      artist: "Neon Harbor",
      title: "Static Horizon",
      votes: 2,
      bpm: 108
    },
    {
      trackId: "trk103",
      artist: "Lunar Arcade",
      title: "Midnight Frequency",
      votes: 4,
      bpm: 128
    }
  ],
  [
    {
      trackId: "trk201",
      artist: "Solar Echo",
      title: "Glass Skyline",
      votes: 3,
      bpm: 115
    },
    {
      trackId: "trk202",
      artist: "Velvet Comet",
      title: "Satellite Hearts",
      votes: 6,
      bpm: 124
    }
  ]
];
```

# --solutions--

```js
function flattenPlaylists(playlists) {
  if (!Array.isArray(playlists)) return [];

  const result = [];

  for (let i = 0; i < playlists.length; i++) {
    const playlist = playlists[i];
    if (!Array.isArray(playlist)) continue;

    for (let j = 0; j < playlist.length; j++) {
      const track = playlist[j];
      result.push({
        trackId: track.trackId,
        artist: track.artist,
        title: track.title,
        votes: track.votes,
        bpm: track.bpm,
        source: [i, j]
      });
    }
  }
  return result;
}

function scoreTracks(tracks) {
  const targetBpm = 120;
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const track = tracks[i];

    result.push({
      trackId: track.trackId,
      artist: track.artist,
      title: track.title,
      votes: track.votes,
      bpm: track.bpm,
      source: track.source,
      score: (track.votes * 10) - Math.abs(track.bpm - targetBpm)
    });
  }
  return result;
}

function dedupeTracks(tracks) {
  const seenTrackIds = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    if (!seenTrackIds.includes(tracks[i].trackId)) {
      seenTrackIds.push(tracks[i].trackId);
      result.push(tracks[i]);
    }
  }
  return result;
}

function enforceArtistQuota(tracks, maxPerArtist) {
  const artists = [];
  const artistCounts = [];
  const result = [];

  for (let i = 0; i < tracks.length; i++) {
    const artist = tracks[i].artist;
    const artistIndex = artists.indexOf(artist);

    if (artistIndex === -1) {
      artists.push(artist);
      artistCounts.push(1);
      result.push(tracks[i]);
    } else if (artistCounts[artistIndex] < maxPerArtist) {
      artistCounts[artistIndex]++;
      result.push(tracks[i]);
    }
  }

  return result;
}

function buildSchedule(tracks) {
  const schedule = [];

  for (let i = 0; i < tracks.length; i++) {
    schedule.push({
      slot: i + 1,
      trackId: tracks[i].trackId
    });
  }

  return schedule;
}

function remixPlaylist(playlists, maxPerArtist) {
  const flattened = flattenPlaylists(playlists);
  const scored = scoreTracks(flattened);
  const deduped = dedupeTracks(scored);
  const limited = enforceArtistQuota(deduped, maxPerArtist);
  return buildSchedule(limited);
}
```
