---
id: 587d78af367417b2b2512b03
title: Crea un Formulario de Encuesta
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. Debe tener un título de página en in elemento `h1` con un `id` de `title`.
1. Debes tener una explicación corta en un elemento `p` con un `id` de `description`.
1. Debes tener un elemento `form` con un `id` de `survey-form`.
1. Dentro del elemento de formulario debes tener un campo `input` requerido para ingresar tu nombre que tiene un `id` de `name` y un `type` de `text`.
1. Dentro del elemento de formulario debes tener un campo `input` requerido para ingresar tu correo electrónico que tiene un `id` de `email`.
1. Si ingresas un correo electrónico que no está formatado correctamente, deberías ver un error de validación HTML5.
1. Dentro del elemento de formulario debes tener un campo `input` para ingresar un número que tiene un `id` de `number`.
1. La entrada de números no debe aceptar caracteres no numéricos, ya sea impidiendo que los escribas o mostrando un error de validación HTML5 (dependiendo del navegador).
1. Si ingresas números fuera del rango del input de número, que están definidos por los atributos `min` y `max`, deberías ver un error de validación HTML5.
1. Para los campos de entrada de nombre, correo electrónico y número, debes tener elementos `label` correspondientes en el formulario que describan el propósito de cada campo con los siguientes ids: `name-label`, `email-label` y `number-label`.
1. Para los campos de entrada de nombre, correo electrónico y número, debes tener un texto de marcador de posición que ofrece una descripción o instrucciones para cada campo.
1. Dentro del elemento de formulario, debes tener un elemento desplegable `select` con un `id` de `dropdown` con al menos dos opciones para elegir.
1. Dentro del elemento de formulario, puedes seleccionar una opción de un grupo de al menos dos botones de radio que son agrupados utilizando el atributo `name`.
1. Dentro del elemento de formulario, puedes seleccionar varios campos en una serie de casillas de verificación, cada una debe tener un atributo `value`.
1. Dentro del elemento de formulario, debes tener un `textarea` para comentarios adicionales.
1. Dentro del elemento de formulario, debes tener un botón con `id` de `submit` para enviar todas las entradas.

# --hints--

Debes tener un elemento `h1` con un `id` de `title`.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

Tu `#title` no debe estar vacío.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Debes tener un elemento `p` con un `id` de `description`.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

Tu `#description` no debe estar vacío.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Debes tener un elemento `form` con un `id` de `survey-form`.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

Debes tener un elemento `input` con un `id` de `name`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Tu `#name` debe tener un `type` de `text`.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

Tu `#name` debe requerir una entrada.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Tu `#name` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

Debes tener un elemento `input` con un `id` de `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Tu `#email` debe tener un `type` de `email`.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

Tu `#email` debe requerir una entrada.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

Tú `#email` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

Debes tener un elemento `input` con un `id` de `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

Tu `#number` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

Tu `#number` debe tener un `type` de `number`.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

Tu `#number` debe tener un atributo `min` con un valor numérico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

Tu `#number` debe tener un atributo `max` con un valor numérico.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

Debes tener un elemento `label` con un `id` de `name-label`.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Debes tener un elemento `label` con un `id` de `email-label`.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Debes tener un elemento `label` con un `id` de `number-label`.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

Tu `#name-label` debe contener texto que describa la entrada.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Tu `#email-label` debe contener texto que describa la entrada.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Tu `#number-label` debe contener texto que describa la entrada.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

Tu `#name-label` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

Tu `#email-label` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

Tu `#number-label` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

Tu `#name` deber tener un atributo `placeholder` y un valor.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Tu `#email` debe tener un atributo `placeholder` y un valor.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Tu `#number` debe tener un atributo `placeholder` y un valor.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

Debes tener un campo `select` con un `id` de `dropdown`.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

Tu `#dropdown` debe tener al menos dos elementos `option` seleccionables (no deshabilitados).

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

Tu `#dropdown` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

Debes tener al menos dos elementos `input` con un `type` de `radio` (botones de radio).

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Debes tener al menos dos botones de radio que sean descendientes de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

Todos los botones de radio deben tener un atributo `value` y un valor.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Todos los botones de radio deben tener un atributo `name` y un valor.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Cada grupo de botones de radio deben tener al menos 2 botones de radio.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

Debes tener al menos dos elementos `input` con un `type` de `checkbox` (casillas de verificación) que sean descendientes de `#survey-form`.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

Todos tus casillas de verificación dentro de `#survey-form` deben tener un atributo `value` y un valor.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

Debes tener al menos un elemento de `textarea` que sea descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

Debes tener un elemento `input` o `button` con un `id` de `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

Tu `#submit` debe tener un `type` de `submit`.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

Tu `#submit` debe ser descendiente de `#survey-form`.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
