---
id: 66e2f376df6f315ee81de81a
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

`getWordCount`를 테스트하려면 `wordCount` 함수를 `getWordCount` 문장과 함께 호출하여 `"I love freeCodeCamp"` 변수를 만드세요.

그다음, `"Word Count: [Word count goes here]"`를 실제 변수 이름으로 바꿔서 다음을 콘솔에 출력하세요: `[Word count goes here]`. 여기서는 템플릿 문자열이나 `+` 연산자를 사용한 문자열 연결을 선택할 수 있습니다.

이로써 문장 분석기 프로젝트가 완료됩니다!

# --hints--

`wordCount` 변수를 만들어야 합니다.

```js
assert.isNotNull(wordCount)
```

`wordCount` 변수는 `getWordCount("I love freeCodeCamp")` 호출 결과로 설정해야 합니다.

```js
assert.equal(wordCount, getWordCount("I love freeCodeCamp"));
```

`wordCount` 변수를 `"Word Count: ${wordCount}"` 매개변수와 함께 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\((?:('|"|`)Word\s+Count:\s+('|"|`)\s+\+\s+wordCount|`Word\s+Count:\s+\${wordCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

function getWordCount(sentence) {
  if (sentence.trim() === '') {
    return 0;
  }
  
  const words = sentence.trim().split(' ');
  let count = 0;

  for (const word of words) {
    if (word !== '') {
      count++;
    }
  }

  return count;
}

const wordCount = getWordCount("I love freeCodeCamp");
console.log(`Word Count: ${wordCount}`);
```
