---
id: 66e2eab8a5638f57b637b7cc
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 문장 부호의 개수를 세세요.

`getPunctuationCount` 함수를 `sentence` 매개변수와 함께 만드세요.

함수 내부에서, 호출될 때 함수에 전달되는 `sentence` 안에 있는 문장 부호 문자의 개수를 세는 반복문을 만드세요. 여기서 문장 부호 문자는 공백(`" "`)이나 글자가 아닌 모든 문자입니다.

`getPunctuationCount` 함수는 숫자를 반환해야 합니다.

# --hints--

`getPunctuationCount` 함수를 만드세요.

```js
assert.isFunction(getPunctuationCount);
```

`getPunctuationCount` 함수는 `sentence` 매개변수를 가져야 합니다.

```js
assert.match(getPunctuationCount.toString(), /sentence/);
```

`getPunctuationCount` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getPunctuationCount("Coding is fun!"))
```

문장이 `"What's going on here?"` 일 때, `getPunctuationCount` 함수는 `2` 을 반환해야 합니다.

```js
assert.strictEqual(getPunctuationCount("What's going on here?"), 2);
```

문장이 `"What????!"` 일 때, `getPunctuationCount` 함수는 `5` 을 반환해야 합니다.

```js
assert.strictEqual(getPunctuationCount("What????!"), 5);
```

`getPunctuationCount` 함수는 어떤 문장에 대해서도 올바른 문장 부호 개수를 반환해야 합니다.

```js
assert.strictEqual(getPunctuationCount("Be quick, sign up! freeCodeCamp awaits, friend!!!"), 6);
assert.strictEqual(getPunctuationCount("Guess what? freeCodeCamp is launching a new cert soon!"), 2);
assert.strictEqual(getPunctuationCount("freeCodeCamp, again? It's incredible!"), 4);
assert.strictEqual(getPunctuationCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
