---
id: 66e2edb23013ca596f8d9568
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 `getPunctuationCount` 함수를 테스트할 시간입니다.

`punctuationCount` 변수를 만들고 `getPunctuationCount` 인수를 사용해 `"WHAT?!?!?!?!?"` 함수를 호출한 결과를 할당하세요.

그다음, 콘솔에 `"Punctuation Count: [Punctuation count goes here]"`를 출력하세요. `[Punctuation count goes here]`를 실제 변수명으로 바꾸세요. 여기서는 템플릿 문자열이나 `+` 연산자를 사용한 문자열 연결 중 선택할 수 있습니다.

# --hints--

`punctuationCount` 변수를 만들어야 합니다.

```js
assert.isNotNull(punctuationCount)
```

`punctuationCount` 변수는 `getPunctuationCount("WHAT?!?!?!?!?")` 결과로 설정해야 합니다.

```js
assert.equal(punctuationCount, getPunctuationCount("WHAT?!?!?!?!?"));
```

콘솔에 `"Punctuation Count: [Punctuation count goes here]"`를 출력해야 합니다. `[Punctuation count goes here]`를 실제 변수명으로 바꾸세요. 문자열 연결 구문을 올바르게 사용했는지 확인하세요.

```js
assert.match(code, /console\.log\((?:('|"|`)Punctuation\s+Count:\s+('|"|`)\s+\+\s+punctuationCount|`Punctuation\s+Count:\s+\${punctuationCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
