---
id: 66e2d680e129e1423116a541
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 문장을 받아 단어 수, 모음, 자음, 구두점 개수를 세는 문장 분석기를 만드세요.

먼저, `getVowelCount` 함수를 만들고 `sentence`이라는 매개변수를 추가하세요. 함수는 문장에서 모음의 총 개수를 반환해야 합니다.

# --hints--

`getVowelCount` 함수를 만들어야 합니다.

```js
assert.isFunction(getVowelCount);
```

`getVowelCount` 함수는 `sentence` 매개변수를 가져야 합니다.

```js
assert.match(getVowelCount.toString(), /sentence/);
```

`getVowelCount` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getVowelCount("Apples are tasty fruits"))
```

문장이 `"Apples are tasty fruits"`일 때, `getVowelCount` 함수는 `7`을 반환해야 합니다.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
```

문장이 `"Hello, World!"`일 때, 함수는 `3`을 반환해야 합니다.

```js
assert.strictEqual(getVowelCount("Hello, World!"), 3);
```

모음 개수는 대소문자를 구분하지 않아야 합니다.

```js
assert.strictEqual(getVowelCount("Apples are tasty fruits"), 7);
assert.strictEqual(getVowelCount("freeCodeCamp is awesome"), 10);
```

`getVowelCount` 함수는 어떤 문장에 대해서도 올바른 모음 개수를 반환해야 합니다.

```js
assert.strictEqual(getVowelCount("I went to the store"), 6);
assert.strictEqual(getVowelCount("The quick brown fox jumps over the lazy dog"), 11);
assert.strictEqual(getVowelCount("The cat in the hat"), 5);
assert.strictEqual(getVowelCount(""), 0);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
