---
id: 66e2e5d890ac8a54c9d703bd
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

자음 개수를 세어 보세요. `getConsonantCount` 함수를 `sentence` 매개변수와 함께 만드세요.

함수 내부에서 반복문을 사용해 함수가 호출될 때 전달받을 `sentence`에 있는 자음 개수를 세세요. 자음은 다음 문자들 `"aeiou"`를 제외한 모든 글자입니다.

`getConsonantCount` 함수는 숫자를 반환해야 합니다.

# --hints--

`getConsonantCount` 함수를 만드세요.

```js
assert.isFunction(getConsonantCount);
```

`getConsonantCount` 함수는 `sentence` 매개변수를 가져야 합니다.

```js
assert.match(getConsonantCount.toString(), /sentence/);
```

`getConsonantCount` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getConsonantCount("Coding is fun"))
```

문장이 `"Coding is fun"`일 때 `getConsonantCount` 함수는 `7`을 반환해야 합니다.

```js
assert.strictEqual(getConsonantCount("Coding is fun"), 7);
```

문장이 `"Hello, World!"`일 때 `getConsonantCount` 함수는 `7`을 반환해야 합니다.

```js
assert.strictEqual(getConsonantCount("Hello, World!"), 7);
```

자음 개수는 대소문자를 구분하지 않아야 합니다.

```js
assert.strictEqual(getConsonantCount("Apples are tasty fruits"), 13);
assert.strictEqual(getConsonantCount("freeCodeCamp is awesome"), 11);
```

`getConsonantCount` 함수는 어떤 문장에 대해서도 올바른 자음 개수를 반환해야 합니다.

```js
assert.strictEqual(getConsonantCount("I went to the store"), 9);
assert.strictEqual(getConsonantCount("The quick brown fox jumps over the lazy dog"), 24);
assert.strictEqual(getConsonantCount("The cat in the hat"), 9);
assert.strictEqual(getConsonantCount(""), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
