---
id: 66e2df8c5676fe51cead8ec8
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 `getVowelCount` 함수를 테스트할 시간입니다.

`vowelCount` 변수를 만들고 `getVowelCount` 인수를 사용해 `"Apples are tasty fruits"` 함수를 호출한 결과를 할당하세요.

그다음, 콘솔에 `"Vowel Count: [vowel count goes here]"`를 출력하세요. `[vowel count goes here]`를 실제 변수 이름으로 바꾸세요. 여기서는 템플릿 문자열이나 `+` 연산자를 사용한 문자열 연결 중 선택할 수 있습니다.

# --hints--

`vowelCount` 변수를 만들어야 합니다.

```js
assert.isNotNull(vowelCount)
```

`vowelCount` 변수는 `getVowelCount("Apples are tasty fruits")` 결과로 설정해야 합니다.

```js
assert.equal(vowelCount, getVowelCount("Apples are tasty fruits"));
```

`vowelCount` 변수를 `"Vowel Count: ${vowelCount}"` 매개변수와 함께 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\((?:('|"|`)Vowel\s+Count:\s+('|"|`)\s+\+\s+vowelCount|`Vowel\s+Count:\s+\${vowelCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
