---
id: 66e2f06b191e305c00574e4d
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

마지막으로, `getWordCount` 함수를 만들어 `sentence` 매개변수를 받으세요. 이 함수는 호출될 때 전달된 문장에 있는 단어의 총 개수를 반환해야 합니다.

# --hints--

`getWordCount` 함수를 만들어야 합니다.

```js
assert.isFunction(getWordCount);
```

`getWordCount` 함수는 `sentence` 매개변수를 가져야 합니다.

```js
assert.match(getWordCount.toString(), /sentence/);
```

`getWordCount` 함수는 숫자를 반환해야 합니다.

```js
assert.isNumber(getWordCount("I love freeCodeCamp"))
```

문장이 `"When are you gonna start learning to code?"` 일 때, `getWordCount` 함수는 `8` 을 반환해야 합니다.

```js
assert.strictEqual(getWordCount("When are you gonna start learning to code?"), 8);
```

문장이 `"What's going on?"` 일 때, `getWordCount` 함수는 `3` 을 반환해야 합니다.

```js
assert.strictEqual(getWordCount("What's going on?"), 3);
```

단어 수는 대소문자를 구분하지 않아야 합니다.

```js
assert.strictEqual(getWordCount("freeCodeCamp offers free coding tutorials online"), 6);
assert.strictEqual(getWordCount("You can learn HTML, CSS, JavaScript, and more on freeCodeCamp"), 10);
```

`getWordCount` 함수는 어떤 문장에 대해서도 올바른 단어 수를 반환해야 합니다.

```js
assert.strictEqual(getWordCount("freeCodeCamp has a great community of kind people"), 8);
assert.strictEqual(getWordCount("The freeCodeCamp curriculum is constantly updated"), 6);
assert.strictEqual(getWordCount("freeCodeCamp teaches both front-end and back-end development"), 7);
```

`getWordCount` 함수는 빈 문자열이나 공백만 있는 문자열에 대해서도 올바른 단어 수를 반환해야 합니다.

```js
assert.strictEqual(getWordCount(""), 0);
assert.strictEqual(getWordCount("   "), 0);
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

const consonantCount = getConsonantCount("Coding is fun");
console.log(`Consonant Count: ${consonantCount}`);

function getPunctuationCount(sentence) {
  const punctuations = ".,!?;:-()[]{}\"'–";
  let count = 0;

  for (const char of sentence) {
    if (punctuations.includes(char)) {
      count++;
    }
  }
  return count;
}

const punctuationCount = getPunctuationCount("WHAT?!?!?!?!?");
console.log(`Punctuation Count: ${punctuationCount}`);

--fcc-editable-region--

--fcc-editable-region--
```
