---
id: 66e2e99caa0cb3570f4822b4
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 `getConsonantCount` 함수를 테스트할 시간입니다.

`consonantCount` 변수를 만들고 `getConsonantCount` 함수를 `"Coding is fun"` 인수와 함께 호출한 결과를 할당하세요.

그다음, 콘솔에 `"Consonant Count: [Consonant count goes here]"`를 출력하세요. `[Consonant count goes here]`를 실제 변수명으로 바꾸세요. 여기서는 템플릿 문자열이나 `+` 연산자를 사용한 문자열 연결 방식을 선택할 수 있습니다.

# --hints--

`consonantCount` 변수를 만들어야 합니다.

```js
assert.isNotNull(consonantCount)
```

`consonantCount` 변수는 `getConsonantCount("Coding is fun")` 결과로 설정해야 합니다.

```js
assert.equal(consonantCount, getConsonantCount("Coding is fun"));
```

`consonantCount` 변수를 `"Consonant Count: ${consonantCount}"` 매개변수와 함께 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\((?:('|"|`)Consonant\s+Count:\s+('|"|`)\s+\+\s+consonantCount|`Consonant\s+Count:\s+\${consonantCount}`)\);?/)
```

# --seed--

## --seed-contents--

```js
function getVowelCount(sentence) {
  const vowels = "aeiou";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (vowels.includes(char)) {
      count++;
    }
  }
  return count;
}

const vowelCount = getVowelCount("Apples are tasty fruits");
console.log(`Vowel Count: ${vowelCount}`);

function getConsonantCount(sentence) {
  const consonants = "bcdfghjklmnpqrstvwxyz";
  let count = 0;

  for (const char of sentence.toLowerCase()) {
    if (consonants.includes(char)) {
      count++;
    }
  }
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
