---
id: 66c06fad3475cd92421b9ac2
title: 创建一个闰年计算器
challengeType: 26
dashedName: build-a-leap-year-calculator
---

# --description--

闰年是能被 `4` 整除的年份，但如果这个年份能被 `100` 整除且不能被 `400` 整除，则不是闰年。 例如，`2000` 年是闰年，但 `1900` 年不是。 另外，闰年在 2 月会多一天，即 2 月有第 29 天。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 定义一个名为 `isLeapYear` 的函数，该函数接受一个数字作为参数。
2. 在函数外部，声明一个变量 `year` 来存储要检查的年份值。
3. 在函数内部，使用 `if/ else` 语句或三元运算符检查年份是否为闰年。
4. 要检查年份是否为闰年，请满足以下条件：

   - 如果年份可被 `4` 整除，则该年为闰年。
   - 除非该年也能被 `100` 整除，则就不是闰年。
   - 除非该年也能被 `400` 整除，则是闰年。

5. 如果年份是闰年，则返回 `[year] is a leap year.`。 否则，返回 `[year] is not a leap year.`。 你将用 `isLeapYear` 函数中定义的参数替换 `[year]`。
6. 你需要调用 `isLeapYear` 函数，并将 `year` 作为参数传入，然后将返回结果赋值给一个名为 `result` 的变量。
7. 你需要使用 `console.log()` 将 `result` 变量输出到控制台。

# --hints--

你应定义一个名为 `isLeapYear` 的函数。

```js
assert.isFunction(isLeapYear);
```

`isLeapYear` 函数应将一个数字作为参数。

```js
assert.match(isLeapYear.toString(), /\s*isLeapYear\(\s*\w+\s*\)/);
```

你应声明一个变量 `year` 并为其赋值，以检查是否为闰年。

```js
assert.isDefined(year);
assert.isNumber(year);
assert.isAtLeast(year, 0);
assert.isTrue(Number.isInteger(year));
```

`year` 变量不应为空。

```js
assert.isNotNull(year);
```

当 `year` 变量的值为 `2024` 时，`result` 应该是 `2024 is a leap year.`

```js
assert.strictEqual(isLeapYear(2024), '2024 is a leap year.');
```

当 `year` 变量的值为 `2000` 时，`result` 应该是 `2000 is a leap year.`

```js
assert.strictEqual(isLeapYear(2000), '2000 is a leap year.');
```

当 `year` 变量的值为 `1900` 时，`result` 应该是 `1900 is not a leap year.`

```js
assert.strictEqual(isLeapYear(1900), '1900 is not a leap year.');

```

你应该调用 `isLeapYear` 函数，并传入 `year` 作为参数。

```js
assert.match(__helpers.removeJSComments(code), /isLeapYear\(\s*year\s*\)/);
```

你应声明一个 `result` 变量。

```js
assert.isDefined(result);
```

你应该把调用 `isLeapYear` 函数的返回结果存储在名为 `result` 的变量中。

```js
assert.match(__helpers.removeJSComments(code), /result\s*=\s*isLeapYear\(\s*year\s*\)/);
```

你应该使用 `console.log()` 将 `result` 输出到控制台。

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function isLeapYear(year) {
  if ((year % 4 === 0 && year % 100 !== 0) || year % 400 === 0) {
    return(year + " is a leap year.");
  } else {
    return(year + " is not a leap year.");
  }
}

const year = 2024;
const result = isLeapYear(year);
console.log(result);
```

