---
id: 56533eb9ac21ba0edf2244b3
title: Tengeneza kivunja joto kutoka Celsius hadi Fahrenheit
challengeType: 26
dashedName: build-a-celsius-to-fahrenheit-converter
---

# --description--

Katika maabara hii, utaandika kitendakazi kinachobadilisha joto kutoka Celsius hadi Fahrenheit. Fomula ya kubadilisha kutoka Celsius hadi Fahrenheit ni:

```js
fahrenheit = celsius * (9/5) + 32
```

**Lengo:** Timiza hadithi za watumizi zilizo hapa chini na upite vipimo vyote ili kumaliza maabara hii.

**Hadithi za watumizi:**

1. Unapaswa kuunda kitendakazi kinachoitwa `convertCtoF`.
2. `convertCtoF` inapaswa kuchukua hoja moja ya nambari, ambayo ni joto katika Celsius.
3. `convertCtoF` inapaswa kurudisha nambari.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `convertCtoF`.

```js
assert.isFunction(convertCtoF);
```

`convertCtoF` inapaswa kuchukua kigezo kimoja.

```js
assert.lengthOf(convertCtoF, 1);
```

`convertCtoF(0)` inapaswa kurudisha nambari.

```js
assert.isNumber(convertCtoF(0));
```

`convertCtoF(-30)` inapaswa kurudisha thamani ya `-22`.

```js
assert.strictEqual(convertCtoF(-30), -22);
```

`convertCtoF(-10)` inapaswa kurudisha thamani ya `14`.

```js
assert.strictEqual(convertCtoF(-10), 14);
```

`convertCtoF(0)` inapaswa kurudisha thamani ya `32`.

```js
assert.strictEqual(convertCtoF(0), 32);
```

`convertCtoF(20)` inapaswa kurudisha thamani ya `68`.

```js
assert.strictEqual(convertCtoF(20), 68);
```

`convertCtoF(30)` inapaswa kurudisha thamani ya `86`.

```js
assert.strictEqual(convertCtoF(30), 86);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function convertCtoF(celsius) {
  let fahrenheit = celsius * (9 / 5) + 32;
  return fahrenheit;
}

convertCtoF(30);
```
