---
id: a26cbbe9ad8655a977e1ceb5
title: 가장 긴 단어 찾기 앱 만들기
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

이 실습에서는 주어진 문장에서 가장 긴 단어의 길이를 반환하는 함수를 만드세요.

예를 들어, 문장 `"The quick brown fox jumped over the lazy dog"`에서 가장 긴 단어는 `"jumped"`이며, 길이는 6입니다.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `findLongestWordLength`라는 이름의 함수를 만들고 문자열을 인수로 받으세요.
2. 함수는 문자열에서 가장 긴 단어의 길이를 반환해야 합니다.

# --hints--

`findLongestWordLength`라는 이름의 함수를 만드세요.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength`는 하나의 매개변수를 가져야 합니다.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")`는 숫자를 반환해야 합니다.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")`는 `6`를 반환해야 합니다.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")`는 `5`를 반환해야 합니다.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")`는 `6`를 반환해야 합니다.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")`는 `8`를 반환해야 합니다.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")`는 `8`를 반환해야 합니다.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")`는 `19`를 반환해야 합니다.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
