---
id: 587d7b8c367417b2b2512b57
title: Використання * для імпорту всього з файлу
challengeType: 1
forumTopicId: 301210
dashedName: use--to-import-everything-from-a-file
---

# --description--

Припустімо, у вас є файл, і ви хочете імпортувати весь його вміст у поточний файл. Це можна зробити за допомогою синтаксису `import * as`. Ось приклад, де вміст файлу з назвою `math_functions.js` імпортується у файл у тій самій директорії:

```js
import * as myMathModule from "./math_functions.js";
```

Вищенаведена інструкція `import` створить об’єкт з назвою `myMathModule`. Це просто ім’я змінної, ви можете назвати її як завгодно. Об’єкт міститиме всі експорти з `math_functions.js`, тож ви можете отримувати доступ до функцій так само, як до будь-якої іншої властивості об’єкта. Ось як можна використати функції `add` і `subtract`, які були імпортовані:

```js
myMathModule.add(2,3);
myMathModule.subtract(5,3);
```

# --instructions--

Код у цьому файлі потребує вміст файлу `string_functions.js`, який знаходиться в тій самій директорії, що й поточний файл. Використайте синтаксис `import * as`, щоб імпортувати все з файлу в об’єкт з назвою `stringFunctions`.

# --hints--

Ваш код має правильно використовувати синтаксис `import * as`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*\*\s*as\s+stringFunctions\s+from\s*('|")\.\/string_functions\.js\1/g
  )
);
```

# --seed--

## --seed-contents--

```js

// Only change code above this line

stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```

# --solutions--

```js
import * as stringFunctions from "./string_functions.js";

// add code above this line
stringFunctions.uppercaseString("hello");
stringFunctions.lowercaseString("WORLD!");
```
