---
id: 587d7b8c367417b2b2512b58
title: Створіть запасний варіант експорту за допомогою export default
challengeType: 1
forumTopicId: 301199
dashedName: create-an-export-fallback-with-export-default
---

# --description--

У `export` уроці ви дізналися про синтаксис, який називають <dfn>іменований експорт</dfn>. Це дозволяло зробити кілька функцій і змінних доступними для використання в інших файлах.

Існує ще один `export` синтаксис, який потрібно знати, відомий як <dfn>export default</dfn>. Зазвичай ви використовуєте цей синтаксис, якщо з файлу експортується лише одне значення. Його також застосовують для створення запасного варіанту значення для файлу або модуля.

Нижче наведено приклади з використанням `export default`:

```js
export default function add(x, y) {
  return x + y;
}

export default function(x, y) {
  return x + y;
}
```

Перший — це іменована функція, а другий — анонімна функція.

Оскільки `export default` використовується для оголошення запасного варіанта значення для модуля або файлу, у кожному модулі або файлі може бути лише одне значення, яке є експортом за замовчуванням. Крім того, ви не можете використовувати `export default` разом із `var`, `let` або `const`.

# --instructions--

Наступна функція має бути запасним варіантом значення для модуля. Будь ласка, додайте необхідний код, щоб це зробити.

# --hints--

Ваш код має використовувати `export` запасний варіант.

```js
assert(
  __helpers.removeJSComments(code).match(
    /export\s+default\s+function(\s+subtract\s*|\s*)\(\s*x,\s*y\s*\)\s*{/g
  )
);
```

# --seed--

## --seed-contents--

```js
function subtract(x, y) {
  return x - y;
}
```

# --solutions--

```js
export default function subtract(x, y) {
  return x - y;
}
```
