---
id: 587d7b88367417b2b2512b46
title: Встановлення параметрів за замовчуванням для ваших функцій
challengeType: 1
forumTopicId: 301209
dashedName: set-default-parameters-for-your-functions
---

# --description--

Щоб допомогти нам створювати більш гнучкі функції, ES6 вводить <dfn>параметри за замовчуванням</dfn> для функцій.

Перегляньте цей код:

```js
const greeting = (name = "Anonymous") => "Hello " + name;

console.log(greeting("John"));
console.log(greeting());
```

У консолі відобразяться рядки `Hello John` та `Hello Anonymous`.

Параметр за замовчуванням спрацьовує, коли аргумент не вказаний (має значення undefined). Як ви бачите у наведеному прикладі, параметр `name` отримає своє значення за замовчуванням `Anonymous`, якщо ви не надасте значення для цього параметра. Ви можете додати значення за замовчуванням для будь-якої кількості параметрів.

# --instructions--

Змініть функцію `increment`, додавши параметри за замовчуванням так, щоб вона додавала 1 до `number`, якщо `value` не вказаний.

# --hints--

Результат `increment(5, 2)` має бути `7`.

```js
assert(increment(5, 2) === 7);
```

Результат `increment(5)` має бути `6`.

```js
assert(increment(5) === 6);
```

Для `1` слід використати значення за замовчуванням `value`.

```js
assert(__helpers.removeJSComments(code).match(/value\s*=\s*1/g));
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const increment = (number, value) => number + value;
// Only change code above this line
```

# --solutions--

```js
const increment = (number, value = 1) => number + value;
```
