---
id: 587d7b89367417b2b2512b4b
title: Використання деструктурованого присвоєння для призначення змінних із масивів
challengeType: 1
forumTopicId: 301213
dashedName: use-destructuring-assignment-to-assign-variables-from-arrays
---

# --description--

ES6 робить деструктуризацію масивів такою ж простою, як і деструктуризацію об’єктів.

Однією з ключових відмінностей між оператором розширення та деструктуризацією масивів є те, що оператор розширення розпаковує весь вміст масиву у список, розділений комами. Відповідно, ви не можете вибирати, які елементи хочете призначити змінним.

Деструктуризація масиву дозволяє зробити саме це:

```js
const [a, b] = [1, 2, 3, 4, 5, 6];
console.log(a, b);
```

У консолі відобразяться значення `a` і `b` як `1, 2`.

Змінній `a` призначене перше значення масиву, а `b` — друге значення масиву. Ми також можемо отримати значення за будь-яким індексом у масиві за допомогою деструктуризації, використовуючи коми, щоб дістатися до потрібного індексу:

```js
const [a, b,,, c] = [1, 2, 3, 4, 5, 6];
console.log(a, b, c);
```

У консолі відобразяться значення `a`, `b` і `c` як `1, 2, 5`.

# --instructions--

Використайте деструктуроване присвоєння, щоб поміняти місцями значення `a` і `b`, так щоб `a` отримала значення, призначене `b`, а `b` — значення, призначене `a`.

# --hints--

Значення `a` має бути `6` після обміну.

```js
assert(a === 6);
```

Значення `b` має бути `8` після обміну.

```js
assert(b === 8);
```

Ви маєте використати деструктуризацію масиву, щоб поміняти місцями `a` і `b`.

```js
assert(/\[\s*(\w)\s*,\s*(\w)\s*\]\s*=\s*\[\s*\2\s*,\s*\1\s*\]/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let a = 8, b = 6;
// Only change code below this line
```

# --solutions--

```js
let a = 8, b = 6;
[a, b] = [b, a];
```
