---
id: 587d7b8d367417b2b2512b59
title: Імпортувати експорт за замовчуванням
challengeType: 1
forumTopicId: 301205
dashedName: import-a-default-export
---

# --description--

У попередньому завданні ви дізналися про `export default` та його застосування. Щоб імпортувати експорт за замовчуванням, потрібно використати інший синтаксис `import`. У наведеному прикладі `add` є експортом за замовчуванням з файлу `math_functions.js`. Ось як його імпортувати:

```js
import add from "./math_functions.js";
```

Синтаксис відрізняється в одному ключовому місці. Імпортоване значення, `add`, не оточене фігурними дужками (`{}`). `add` тут — це просто ім'я змінної для того, що є експортом за замовчуванням з файлу `math_functions.js`. Ви можете використовувати будь-яке ім'я при імпорті за замовчуванням.

# --instructions--

У наведеному коді імпортуйте експорт за замовчуванням з файлу `math_functions.js`, який знаходиться в тій самій директорії, що й цей файл. Дайте імпорту ім'я `subtract`.

# --hints--

Ви повинні правильно імпортувати `subtract` з `math_functions.js`.

```js
assert(__helpers.removeJSComments(code).match(/import\s+subtract\s+from\s+('|")\.\/math_functions\.js\1/g));
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

subtract(7,4);
```

# --solutions--

```js
import subtract from "./math_functions.js";

subtract(7,4);
```
