---
id: 587d7b8c367417b2b2512b54
title: Використання геттерів і сеттерів для контролю доступу до об’єкта
challengeType: 1
forumTopicId: 301220
dashedName: use-getters-and-setters-to-control-access-to-an-object
---

# --description--

Ви можете отримувати значення з об’єкта та встановлювати значення властивості всередині об’єкта.

Ці методи класично називають <dfn>геттерами</dfn> та <dfn>сеттерами</dfn>.

Геттер-функції призначені для того, щоб просто повертати (отримувати) значення приватної змінної об’єкта користувачу без прямого доступу користувача до цієї приватної змінної.

Сеттер-функції призначені для зміни (встановлення) значення приватної змінної об’єкта на основі значення, переданого в сеттер-функцію. Це може включати обчислення або навіть повне перезаписування попереднього значення.

```js
class Book {
  constructor(author) {
    this._author = author;
  }
  // getter
  get writer() {
    return this._author;
  }
  // setter
  set writer(updatedAuthor) {
    this._author = updatedAuthor;
  }
}
const novel = new Book('anonymous');
console.log(novel.writer);
novel.writer = 'newAuthor';
console.log(novel.writer);
```

У консолі відобразяться рядки `anonymous` та `newAuthor`.

Зверніть увагу на синтаксис, який використовується для виклику геттера і сеттера. Вони навіть не виглядають як функції. Геттери і сеттери важливі, бо приховують внутрішні деталі реалізації.

**Примітка:** За конвенцією ім’я приватної змінної починають з підкреслення (`_`). Однак сама ця практика не робить змінну приватною.

# --instructions--

Використайте ключове слово `class` для створення класу `Thermostat`. Конструктор `constructor` приймає температуру у Фаренгейтах.

У класі створіть геттер `getter` для отримання температури в Цельсіях і сеттер `setter`, який приймає температуру в Цельсіях.

Пам’ятайте, що `C = 5/9 * (F - 32)` і `F = C * 9.0 / 5 + 32`, де `F` — це значення температури у Фаренгейтах, а `C` — значення тієї ж температури у Цельсіях.

**Примітка:** Коли ви реалізуєте це, температура всередині класу буде зберігатися в одній шкалі — або у Фаренгейтах, або в Цельсіях.

Ось у чому сила геттера і сеттера. Ви створюєте API для іншого користувача, який може отримати правильний результат незалежно від того, в якій шкалі ви зберігаєте температуру.

Інакше кажучи, ви абстрагуєте деталі реалізації від користувача.

# --hints--

`Thermostat` має бути `class` з визначеним методом `constructor`.

```js
assert.isFunction(Thermostat);
assert.isFunction(Thermostat?.constructor);
```

Потрібно використати ключове слово `class`.

```js
assert.match(code, /class/);
```

`Thermostat` має бути можливим для створення екземпляра.

```js
const _t = new Thermostat(122);
assert.isObject(_t);
```

При створенні екземпляра з температурою у Фаренгейтах `Thermostat` має встановити правильне `temperature`.

```js
const _t = new Thermostat(122);
assert.strictEqual(_t?.temperature, 50);
```

Потрібно визначити геттер `getter`.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.get);
```

Потрібно визначити сеттер `setter`.

```js
const _desc = Object.getOwnPropertyDescriptor(Thermostat.prototype, 'temperature');
assert.isFunction(_desc?.set);
```

Виклик сеттера `setter` із температурою в Цельсіях має встановити `temperature`.

```js
const _t = new Thermostat(32);
_t.temperature = 26;
const _u = new Thermostat(32);
_u.temperature = 50;
assert.approximately(_t.temperature, 26, 0.1);
assert.approximately(_u.temperature, 50, 0.1);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```

# --solutions--

```js
class Thermostat {
  constructor(fahrenheit) {
    this._tempInCelsius = 5/9 * (fahrenheit - 32);
  }
  get temperature(){
    return this._tempInCelsius;
  }
  set temperature(newTemp){
    this._tempInCelsius = newTemp;
  }
}

const thermos = new Thermostat(76); // Setting in Fahrenheit scale
let temp = thermos.temperature; // 24.44 in Celsius
thermos.temperature = 26;
temp = thermos.temperature; // 26 in Celsius
```
