---
id: 587d7b88367417b2b2512b47
title: Використання параметра залишку з параметрами функції
challengeType: 1
forumTopicId: 301221
dashedName: use-the-rest-parameter-with-function-parameters
---

# --description--

Щоб допомогти нам створювати більш гнучкі функції, ES6 вводить <dfn>параметр залишку</dfn> для параметрів функції. За допомогою параметра залишку ви можете створювати функції, які приймають змінну кількість аргументів. Ці аргументи зберігаються в масиві, до якого можна отримати доступ пізніше всередині функції.

Перегляньте цей код:

```js
function howMany(...args) {
  return "You have passed " + args.length + " arguments.";
}
console.log(howMany(0, 1, 2));
console.log(howMany("string", null, [1, 2, 3], { }));
```

У консолі відобразяться рядки `You have passed 3 arguments.` та `You have passed 4 arguments.`.

Параметр залишку усуває необхідність використовувати об’єкт `arguments` і дозволяє нам використовувати методи масиву для масиву параметрів, переданих у функцію `howMany`.

# --instructions--

Змініть функцію `sum`, використовуючи параметр залишку так, щоб функція `sum` могла приймати будь-яку кількість аргументів і повертати їх суму.

# --hints--

Результат `sum(0,1,2)` має бути 3

```js
assert(sum(0, 1, 2) === 3);
```

Результат `sum(1,2,3,4)` має бути 10

```js
assert(sum(1, 2, 3, 4) === 10);
```

Результат `sum(5)` має бути 5

```js
assert(sum(5) === 5);
```

Результат `sum()` має бути 0

```js
assert(sum() === 0);
```

`sum` має бути стрілковою функцією, яка використовує синтаксис параметра залишку (`...`) для параметра `args`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/sum=\(\.\.\.args\)=>/));
```

# --seed--

## --seed-contents--

```js
const sum = (x, y, z) => {
  const args = [x, y, z];
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```

# --solutions--

```js
const sum = (...args) => {
  let total = 0;
  for (let i = 0; i < args.length; i++) {
    total += args[i];
  }
  return total;
}
```
