---
id: 587d7b89367417b2b2512b4a
title: Використання деструктурованого присвоєння для призначення змінних зі вкладених об’єктів
challengeType: 1
forumTopicId: 301214
dashedName: use-destructuring-assignment-to-assign-variables-from-nested-objects
---

# --description--

Ви можете застосувати ті самі принципи з попередніх двох уроків, щоб деструктурувати значення зі вкладених об’єктів.

Використовуючи об’єкт, подібний до попередніх прикладів:

```js
const user = {
  johnDoe: { 
    age: 34,
    email: 'johnDoe@freeCodeCamp.com'
  }
};
```

Ось як можна витягти значення властивостей об’єкта та призначити їх змінним з такими ж іменами:

```js
const { johnDoe: { age, email }} = user;
```

А ось як можна призначити значення властивостей об’єкта змінним з іншими іменами:

```js
const { johnDoe: { age: userAge, email: userEmail }} = user;
```

# --instructions--

Замініть два присвоєння еквівалентним деструктурованим присвоєнням. Воно має призначити змінним `lowToday` і `highToday` значення `today.low` і `today.high` з об’єкта `LOCAL_FORECAST`.

# --hints--

Вам слід видалити синтаксис присвоєння ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/lowToday = LOCAL_FORECAST\.today\.low/g) &&
    !__helpers.removeJSComments(code).match(/highToday = LOCAL_FORECAST\.today.high/g)
);
```

Використайте деструктуризацію, щоб створити змінну `lowToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(low\s*:\s*lowToday[^}]*|[^,]*,\s*low\s*:\s*lowToday\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

Використайте деструктуризацію, щоб створити змінну `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*today\s*:\s*{\s*(high\s*:\s*highToday[^}]*|[^,]*,\s*high\s*:\s*highToday,?\s*)},?\s*}\s*=\s*LOCAL_FORECAST(;|\s+|\/\/)/g
  )
);
```

`lowToday` має бути рівним `64`, а `highToday` має бути рівним `77`.

```js
assert(lowToday === 64 && highToday === 77);
```

# --seed--

## --seed-contents--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};

// Only change code below this line
  
const lowToday = LOCAL_FORECAST.today.low;
const highToday = LOCAL_FORECAST.today.high;

// Only change code above this line
```

# --solutions--

```js
const LOCAL_FORECAST = {
  yesterday: { low: 61, high: 75 },
  today: { low: 64, high: 77 },
  tomorrow: { low: 68, high: 80 }
};
 
const { today: { low: lowToday, high: highToday }} = LOCAL_FORECAST;
```
