---
id: 587d7b89367417b2b2512b48
title: Використання оператора розширення для оцінки масивів на місці
challengeType: 1
forumTopicId: 301222
dashedName: use-the-spread-operator-to-evaluate-arrays-in-place
---

# --description--

ES6 вводить <dfn>оператор розширення</dfn>, який дозволяє розгортати масиви та інші вирази в тих місцях, де очікується кілька параметрів або елементів.

Код ES5 нижче використовує `apply()` для обчислення максимального значення в масиві:

```js
var arr = [6, 89, 3, 45];
var maximus = Math.max.apply(null, arr);
```

`maximus` матиме значення `89`.

Нам довелося використати `Math.max.apply(null, arr)`, тому що `Math.max(arr)` повертає `NaN`. `Math.max()` очікує аргументи, розділені комами, а не масив. Оператор розширення робить цей синтаксис набагато зручнішим для читання та підтримки.

```js
const arr = [6, 89, 3, 45];
const maximus = Math.max(...arr);
```

`maximus` матиме значення `89`.

`...arr` повертає розгорнутий масив. Іншими словами, він розгортає масив. Проте оператор розширення працює лише на місці, наприклад, як аргумент функції або в літералі масиву. Наприклад:

```js
const spreaded = [...arr];
```

Однак наступний код не працюватиме:

```js
const spreaded = ...arr;
```

# --instructions--

Скопіюйте весь вміст `arr1` в інший масив `arr2` за допомогою оператора розширення.

# --hints--

`arr2` має бути правильною копією `arr1`.

```js
assert(arr2.every((v, i) => v === arr1[i]) && arr2.length);
```

Для дублювання `...` слід використати оператор розширення `arr1`.

```js
assert(__helpers.removeJSComments(code).match(/Array\(\s*\.\.\.arr1\s*\)|\[\s*\.\.\.arr1\s*\]/));
```

`arr2` має залишатися незмінним, коли змінюють `arr1`.

```js
assert((arr1, arr2) => {
  arr1.push('JUN');
  return arr2.length < arr1.length;
});
```

# --seed--

## --seed-contents--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [];  // Change this line

console.log(arr2);
```

# --solutions--

```js
const arr1 = ['JAN', 'FEB', 'MAR', 'APR', 'MAY'];
let arr2;

arr2 = [...arr1];
```
