---
id: 587d7b8b367417b2b2512b53
title: Використання синтаксису класу для визначення конструктора
challengeType: 1
forumTopicId: 301212
dashedName: use-class-syntax-to-define-a-constructor-function
---

# --description--

ES6 надає новий синтаксис для створення об’єктів за допомогою ключового слова <dfn>class</dfn>.

У ES5 об’єкт можна створити, оголосивши `constructor` функцію та використавши ключове слово `new` для створення екземпляра об’єкта.

У ES6 оголошення `class` має метод `constructor`, який викликається за допомогою ключового слова `new`. Якщо метод `constructor` явно не визначений, він неявно визначається без аргументів.

```js
// Explicit constructor
class SpaceShuttle {
  constructor(targetPlanet) {
    this.targetPlanet = targetPlanet;
  }
  takeOff() {
    console.log("To " + this.targetPlanet + "!");
  }
}

// Implicit constructor 
class Rocket {
  launch() {
    console.log("To the moon!");
  }
}

const zeus = new SpaceShuttle('Jupiter');
// prints To Jupiter! in console
zeus.takeOff();

const atlas = new Rocket();
// prints To the moon! in console
atlas.launch();
```

Слід зазначити, що ключове слово `class` оголошує нову функцію, до якої додається конструктор. Цей конструктор викликається, коли виконується `new` для створення нового об’єкта.

**Примітка:** За конвенцією для імен класів ES6 слід використовувати UpperCamelCase, як у прикладі `SpaceShuttle` вище.

Метод `constructor` — це спеціальний метод для створення та ініціалізації об’єкта, створеного за допомогою класу. Ви дізнаєтеся про нього більше в розділі об’єктно-орієнтованого програмування сертифікації JavaScript Algorithms And Data Structures.

# --instructions--

Використайте ключове слово `class` і напишіть `constructor` для створення класу `Vegetable`.

Клас `Vegetable` дозволяє створити об’єкт овоча з властивістю `name`, яка передається до `constructor`.

# --hints--

`Vegetable` має бути `class` з визначеним методом `constructor`.

```js
assert(
  typeof Vegetable === 'function' && typeof Vegetable.constructor === 'function'
);
```

Слід використати ключове слово `class`.

```js
assert(__helpers.removeJSComments(code).match(/class/g));
```

`Vegetable` має бути можливим для створення екземпляра.

```js
assert(() => {
  const a = new Vegetable('apple');
  return typeof a === 'object';
});
```

`carrot.name` має повертати `carrot`.

```js
assert(carrot.name == 'carrot');
```

# --seed--

## --seed-contents--

```js
// Only change code below this line

// Only change code above this line

const carrot = new Vegetable('carrot');
console.log(carrot.name); // Should display 'carrot'
```

# --solutions--

```js
class Vegetable {
  constructor(name) {
    this.name = name;
  }
}
const carrot = new Vegetable('carrot');
```
