---
id: 587d7b8a367417b2b2512b4f
title: Напишіть лаконічні оголошення літералів об’єктів за допомогою скорочення властивостей об’єкта
challengeType: 1
forumTopicId: 301225
dashedName: write-concise-object-literal-declarations-using-object-property-shorthand
---

# --description--

ES6 додає зручну підтримку для легкого визначення літералів об’єктів.

Розгляньте наступний код:

```js
const getMousePosition = (x, y) => ({
  x: x,
  y: y
});
```

`getMousePosition` — це проста функція, яка повертає об’єкт із двома властивостями. ES6 надає синтаксичний цукор, щоб усунути надмірність необхідності писати `x: x`. Ви можете просто написати `x` один раз, і він буде перетворений на `x: x` (або щось еквівалентне) під капотом. Ось та сама функція з вище, переписана з використанням цього нового синтаксису:

```js
const getMousePosition = (x, y) => ({ x, y });
```

# --instructions--

Використайте скорочення властивостей об’єкта з літералами об’єктів, щоб створити та повернути об’єкт із властивостями `name`, `age` і `gender`.

# --hints--

`createPerson("Zodiac Hasbro", 56, "male")` має повертати `{name: "Zodiac Hasbro", age: 56, gender: "male"}`.

```js
assert.deepEqual(
  { name: 'Zodiac Hasbro', age: 56, gender: 'male' },
  createPerson('Zodiac Hasbro', 56, 'male')
);
```

Ваш код не повинен використовувати `key:value`.

```js
assert(!__helpers.removeJSComments(code).match(/:/g))
```

# --seed--

## --seed-contents--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  return {
    name: name,
    age: age,
    gender: gender
  };
  // Only change code above this line
};
```

# --solutions--

```js
const createPerson = (name, age, gender) => {
  // Only change code below this line
  /*return {
    name: name,
    age: age,
    gender: gender
  };*/
  return {
    name,
    age,
    gender
  };
  // Only change code above this line
};
```
