---
id: 587d7b87367417b2b2512b43
title: Використовуйте стрілкові функції для написання лаконічних анонімних функцій
challengeType: 1
forumTopicId: 301211
dashedName: use-arrow-functions-to-write-concise-anonymous-functions
---

# --description--

У JavaScript часто не потрібно давати ім’я нашим функціям, особливо коли передаємо функцію як аргумент іншій функції. Замість цього ми створюємо рядкові функції. Їх не потрібно іменувати, бо ми не використовуємо їх повторно в інших місцях.

Для цього ми часто використовуємо такий синтаксис:

```js
const myFunc = function() {
  const myVar = "value";
  return myVar;
}
```

ES6 надає нам синтаксичний цукор, щоб не писати анонімні функції таким чином. Замість цього можна використати **синтаксис стрілкової функції**:

```js
const myFunc = () => {
  const myVar = "value";
  return myVar;
}
```

Коли немає тіла функції, а лише повернене значення, синтаксис стрілкової функції дозволяє опустити ключове слово `return` і дужки навколо коду. Це допомагає спростити невеликі функції до однорядкових інструкцій:

```js
const myFunc = () => "value";
```

Цей код за замовчуванням все одно поверне рядок `value`.

# --instructions--

Перепишіть функцію, призначену змінній `magic`, яка повертає `new Date()`, використовуючи синтаксис стрілкової функції. Також переконайтеся, що нічого не визначено за допомогою ключового слова `var`.

# --hints--

Вам слід замінити ключове слово `var`.

```js
assert.notMatch(code, /var/g)
```

`magic` має бути константою (за допомогою `const`).

```js
assert.match(code, /const\s+magic/g)
```

`magic` має бути `function`.

```js
assert(typeof magic === 'function');
```

`magic()` має повертати правильну дату.

```js
assert(magic().setHours(0, 0, 0, 0) === new Date().setHours(0, 0, 0, 0));
```

Ключове слово `function` не повинно використовуватися.

```js
assert.notMatch(code, /function/g)
```

# --seed--

## --seed-contents--

```js
var magic = function() {
  return new Date();
};
```

# --solutions--

```js
const magic = () => {
  return new Date();
};
```
