---
id: 587d7b89367417b2b2512b49
title: Використання деструктурованого присвоєння для призначення змінних з об’єктів
challengeType: 1
forumTopicId: 301215
dashedName: use-destructuring-assignment-to-assign-variables-from-objects
---

# --description--

Деструктуризація дозволяє призначати нову назву змінній під час вилучення значень. Ви можете зробити це, поставивши нову назву після двокрапки під час присвоєння значення.

Використовуючи той самий об’єкт з попереднього прикладу:

```js
const user = { name: 'John Doe', age: 34 };
```

Ось як ви можете призначити нові назви змінним у присвоєнні:

```js
const { name: userName, age: userAge } = user;
```

Ви можете прочитати це як «отримати значення `user.name` і призначити його новій змінній з назвою `userName`» і так далі. Значення `userName` буде рядком `John Doe`, а значення `userAge` буде числом `34`.

# --instructions--

Замініть два присвоєння еквівалентним деструктурованим присвоєнням. Воно має призначити змінним `highToday` і `highTomorrow` значення `today` і `tomorrow` з об’єкта `HIGH_TEMPERATURES`.

# --hints--

Вам слід видалити синтаксис присвоєння ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/highToday = HIGH_TEMPERATURES\.today/g) &&
    !__helpers.removeJSComments(code).match(/highTomorrow = HIGH_TEMPERATURES\.tomorrow/g)
);
```

Використайте деструктуризацію, щоб створити змінну `highToday`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(today\s*:\s*highToday[^}]*|[^,]*,\s*today\s*:\s*highToday\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

Використайте деструктуризацію, щоб створити змінну `highTomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(var|const|let)\s*{\s*(tomorrow\s*:\s*highTomorrow[^}]*|[^,]*,\s*tomorrow\s*:\s*highTomorrow\s*)}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g
  )
);
```

`highToday` має бути рівним `77`, а `highTomorrow` має бути рівним `80`.

```js
assert(highToday === 77 && highTomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line
  
const highToday = HIGH_TEMPERATURES.today;
const highTomorrow = HIGH_TEMPERATURES.tomorrow; 

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today: highToday, tomorrow: highTomorrow } = HIGH_TEMPERATURES;
```
