---
id: 587d7b8c367417b2b2512b55
title: Повторне використання коду JavaScript за допомогою import
challengeType: 1
forumTopicId: 301208
dashedName: reuse-javascript-code-using-import
---

# --description--

`import` дозволяє вибирати, які частини файлу або модуля завантажувати. У попередньому уроці приклади експортували `add` з файлу `math_functions.js`. Ось як ви можете імпортувати його для використання в іншому файлі:

```js
import { add } from './math_functions.js';
```

Тут `import` знайде `add` у `math_functions.js`, імпортує лише цю функцію для вашого використання і проігнорує решту. `./` вказує імпорту шукати файл `math_functions.js` у тій же папці, що й поточний файл. Відносний шлях до файлу (`./`) і розширення файлу (`.js`) потрібні при використанні import таким чином.

Ви можете імпортувати більше одного елемента з файлу, додавши їх у `import` інструкцію ось так:

```js
import { add, subtract } from './math_functions.js';
```

# --instructions--

Додайте відповідну `import` інструкцію, яка дозволить поточному файлу використовувати функції `uppercaseString` і `lowercaseString`, які ви експортували в попередньому уроці. Ці функції знаходяться у файлі `string_functions.js`, який розташований у тій же директорії, що й поточний файл.

# --hints--

Ви повинні правильно імпортувати `uppercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

Ви повинні правильно імпортувати `lowercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /import\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)}\s+from\s+('|")\.\/string_functions\.js\2/g
  )
);
```

# --seed--

## --seed-contents--

```js
  
// Only change code above this line

uppercaseString("hello");
lowercaseString("WORLD!");
```

# --solutions--

```js
import { uppercaseString, lowercaseString } from './string_functions.js';

uppercaseString("hello");
lowercaseString("WORLD!");
```
