---
id: 587d7b8a367417b2b2512b4c
title: >-
  Деструктуроване присвоєння через елементи rest
challengeType: 1
forumTopicId: 301218
dashedName: >-
  use-destructuring-assignment-with-the-rest-parameter-to-reassign-array-elements
---

# --description--

У деяких випадках деструктуризації масивів ми можемо захотіти зібрати залишок елементів у окремий масив.

Результат схожий на `Array.prototype.slice()`, як показано нижче:

```js
const [a, b, ...arr] = [1, 2, 3, 4, 5, 7];
console.log(a, b);
console.log(arr);
```

У консолі відобразяться значення `1, 2` і `[3, 4, 5, 7]`.

Змінні `a` і `b` отримують перше та друге значення з масиву. Після цього, через наявність синтаксису rest, `arr` отримує залишок значень у вигляді масиву. Елемент rest працює коректно лише як остання змінна у списку. Тобто, ви не можете використати синтаксис rest, щоб захопити підмасив, який пропускає останній елемент оригінального масиву.

# --instructions--

Використайте деструктуроване присвоєння з синтаксисом rest, щоб імітувати поведінку `Array.prototype.slice()`. `removeFirstTwo()` має повертати підмасив оригінального масиву `list`, пропускаючи перші два елементи.

# --hints--

`removeFirstTwo([1, 2, 3, 4, 5])` має бути `[3, 4, 5]`

```js
assert.deepEqual(removeFirstTwo([1, 2, 3, 4, 5]), [3, 4, 5]);
```

`removeFirstTwo()` не повинен змінювати `list`

```js
const _testArr = [1, 2, 3, 4, 5];
removeFirstTwo(_testArr);
assert.deepEqual(_testArr, [1, 2, 3, 4, 5])
```

`Array.slice()` не слід використовувати.

```js
assert(!__helpers.removeJSComments(code).match(/\.\s*slice\s*\(/));
```

Вам слід використати синтаксис rest.

```js
assert.match(code, /\.\.\./);
```

# --seed--

## --seed-contents--

```js
function removeFirstTwo(list) {
  return list;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```

# --solutions--

```js
function removeFirstTwo(list) {
  // comment with 'slice' to check comments are removed in tests
  const [, , ...shorterList] = list;
  return shorterList;
}

const source = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const sourceWithoutFirstTwo = removeFirstTwo(source);
```
