---
id: 587d7b8a367417b2b2512b4d
title: "Використовуйте деструктуроване присвоєння, щоб передати об’єкт як параметри функції"
challengeType: 1
forumTopicId: 301217
dashedName: use-destructuring-assignment-to-pass-an-object-as-a-functions-parameters
---

# --description--

У деяких випадках можна деструктурувати об’єкт безпосередньо в аргументі функції.

Розгляньте код нижче:

```js
const profileUpdate = (profileData) => {
  const { name, age, nationality, location } = profileData;

}
```

Це фактично деструктурує об’єкт, переданий у функцію. Це також можна зробити безпосередньо на місці:

```js
const profileUpdate = ({ name, age, nationality, location }) => {

}
```

Коли `profileData` передається у вказану функцію, значення деструктуруються з параметра функції для використання всередині функції.

# --instructions--

Використайте деструктуроване присвоєння в аргументі функції `half`, щоб передати всередину функції лише `max` і `min`.

# --hints--

`stats` має бути `object`.

```js
assert(typeof stats === 'object');
```

`half(stats)` має бути `28.015`

```js
assert(half(stats) === 28.015);
```

Потрібно використати деструктуризацію.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/half=\({\w+,\w+}\)/));
```

Потрібно використати деструктурований параметр.

```js
assert(!__helpers.removeJSComments(code).match(/stats\.max|stats\.min/));
```

# --seed--

## --seed-contents--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

// Only change code below this line
const half = (stats) => (stats.max + stats.min) / 2.0; 
// Only change code above this line
```

# --solutions--

```js
const stats = {
  max: 56.78,
  standard_deviation: 4.34,
  median: 34.54,
  mode: 23.87,
  min: -0.75,
  average: 35.85
};

const half = ( {max, min} ) => (max + min) / 2.0;
```
