---
id: 587d7b8b367417b2b2512b50
title: Написання лаконічних декларативних функцій з ES6
challengeType: 1
forumTopicId: 301224
dashedName: write-concise-declarative-functions-with-es6
---

# --description--

При визначенні функцій у межах об’єктів у ES5 потрібно використовувати ключове слово `function` таким чином:

```js
const person = {
  name: "Taylor",
  sayHello: function() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

У ES6 можна повністю прибрати ключове слово `function` і двокрапку при визначенні функцій у об’єктах. Ось приклад такого синтаксису:

```js
const person = {
  name: "Taylor",
  sayHello() {
    return `Hello! My name is ${this.name}.`;
  }
};
```

# --instructions--

Перепишіть функцію `setGear` всередині об’єкта `bicycle`, щоб вона використовувала скорочений синтаксис, описаний вище.

# --hints--

Не слід використовувати традиційний функціональний вираз.

```js
assert(!__helpers.removeJSComments(code).match(/function/));
```

`setGear` має бути декларативною функцією.

```js
assert(
  typeof bicycle.setGear === 'function' && __helpers.removeJSComments(code).match(/setGear\s*\(.+\)\s*\{/)
);
```

`bicycle.setGear(48)` має змінити значення `gear` на 48.

```js
bicycle.setGear(48);
assert(bicycle.gear === 48);
```

# --seed--

## --seed-contents--

```js
// Only change code below this line
const bicycle = {
  gear: 2,
  setGear: function(newGear) {
    this.gear = newGear;
  }
};
// Only change code above this line
bicycle.setGear(3);
console.log(bicycle.gear);
```

# --solutions--

```js
const bicycle = {
  gear: 2,
  // setGear: function(newGear) {
  setGear(newGear) {
    this.gear = newGear;
  }
};
bicycle.setGear(3);
```
