---
id: 587d7b88367417b2b2512b44
title: Написання стрілкових функцій із параметрами
challengeType: 1
forumTopicId: 301223
dashedName: write-arrow-functions-with-parameters
---

# --description--

Так само, як і звичайна функція, ви можете передавати аргументи у стрілкову функцію.

```js
const doubler = (item) => item * 2;
doubler(4);
```

`doubler(4)` поверне значення `8`.

Якщо стрілкова функція має один параметр, круглі дужки навколо параметра можна опустити.

```js
const doubler = item => item * 2;
```

Можна передавати більше одного аргумента у стрілкову функцію.

```js
const multiplier = (item, multi) => item * multi;
multiplier(4, 2);
```

`multiplier(4, 2)` поверне значення `8`.

# --instructions--

Перепишіть функцію `myConcat`, яка додає вміст `arr2` до `arr1`, так щоб функція використовувала синтаксис стрілкової функції.

# --hints--

Вам слід замінити ключове слово `var`.

```js
assert.notMatch(code, /var/g);
```

`myConcat` має бути константою (за допомогою `const`).

```js
assert.match(code, /const\s+myConcat/g);
```

`myConcat` має бути стрілковою функцією з двома параметрами.

```js
assert(
  /myConcat=\(\w+,\w+\)=>/.test(__helpers.removeJSComments(code).replace(/\s/g, '')) &&
    typeof myConcat === 'function'
);
```

`myConcat()` має повертати `[1, 2, 3, 4, 5]`.

```js
assert.deepEqual(myConcat([1, 2], [3, 4, 5]), [1, 2, 3, 4, 5]);
```

Ключове слово `function` не повинно використовуватися.

```js
assert.notMatch(code, /function/g);
```

# --seed--

## --seed-contents--

```js
var myConcat = function(arr1, arr2) {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```

# --solutions--

```js
const myConcat = (arr1, arr2) =>  {
  return arr1.concat(arr2);
};

console.log(myConcat([1, 2], [3, 4, 5]));
```
