---
id: 587d7b8a367417b2b2512b4e
title: Створення рядків за допомогою шаблонних літералів
challengeType: 1
forumTopicId: 301200
dashedName: create-strings-using-template-literals
---

# --description--

Нова функція ES6 — це <dfn>шаблонний літерал</dfn>. Це особливий тип рядка, який полегшує створення складних рядків.

Шаблонні літерали дозволяють створювати багаторядкові рядки та використовувати інтерполяцію рядків для їх формування.

Розгляньте код нижче:

```js
const person = {
  name: "Zodiac Hasbro",
  age: 56
};

const greeting = `Hello, my name is ${person.name}!
I am ${person.age} years old.`;

console.log(greeting);
```

У консолі відобразяться рядки `Hello, my name is Zodiac Hasbro!` та `I am 56 years old.`.

Там сталося багато чого. По-перше, у прикладі використовуються зворотні апострофи (`` ` ``), not quotes (`'` or `"`), to wrap the string. Secondly, notice that the string is multi-line, both in the code and the output. This saves inserting `\n` within strings. The `${variable}` syntax used above is a placeholder. Basically, you won't have to use concatenation with the `+` operator anymore. To add variables to strings, you just drop the variable in a template string and wrap it with `${` and `}`. Similarly, you can include other expressions in your string literal, for example `${a + b}`). Цей новий спосіб створення рядків дає більше гнучкості для формування надійних рядків.

# --instructions--

Використайте синтаксис шаблонних літералів із зворотними апострофами, щоб створити масив рядків елементів списку (`li`). Текст кожного елемента списку має бути одним із елементів масиву з властивості `failure` об’єкта `result` і мати атрибут `class` зі значенням `text-warning`. Функція `makeList` має повертати масив рядків елементів списку.

Використайте ітератор (будь-який тип циклу), щоб отримати потрібний вихідний результат (показаний нижче).

```js
[
  '<li class="text-warning">no-var</li>',
  '<li class="text-warning">var-on-top</li>',
  '<li class="text-warning">linebreak</li>'
]
```

# --hints--

`failuresList` має бути масивом, що містить `result failure` повідомлень.

```js
assert(
  typeof makeList(result.failure) === 'object' && failuresList.length === 3
);
```

`failuresList` має дорівнювати вказаному вихідному результату.

```js
assert(
  makeList(result.failure).every(
    (v, i) =>
      v === `<li class="text-warning">${result.failure[i]}</li>` ||
      v === `<li class='text-warning'>${result.failure[i]}</li>`
  )
);
```

Слід використовувати шаблонні рядки та інтерполяцію виразів.

```js
assert.match(code, /(`.*\${.*}.*`)/);
```

Потрібно використати ітератор.

```js
assert(__helpers.removeJSComments(code).match(/for|map|reduce|forEach|while/));
```

# --seed--

## --seed-contents--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  // Only change code below this line
  const failureItems = [];
  // Only change code above this line

  return failureItems;
}

const failuresList = makeList(result.failure);
```

# --solutions--

```js
const result = {
  success: ["max-length", "no-amd", "prefer-arrow-functions"],
  failure: ["no-var", "var-on-top", "linebreak"],
  skipped: ["no-extra-semi", "no-dup-keys"]
};
function makeList(arr) {
  return arr.map(val => `<li class="text-warning">${val}</li>`);
}

const failuresList = makeList(result.failure);
```
