---
id: 587d7b87367417b2b2512b42
title: "Мутувати масив, оголошений за допомогою const"
challengeType: 1
forumTopicId: 301206
dashedName: mutate-an-array-declared-with-const
---

# --description--

Якщо ви не знайомі з `const`, ознайомтеся з <a href="/learn/javascript-algorithms-and-data-structures/basic-javascript/declare-a-read-only-variable-with-the-const-keyword" target="_blank" rel="noopener noreferrer nofollow">цим завданням про ключове слово <code>const</code></a>.

Оголошення `const` має багато випадків використання в сучасному JavaScript.

Деякі розробники віддають перевагу призначати всі свої змінні за допомогою `const` за замовчуванням, якщо вони не знають, що їм доведеться переназначити значення. Лише в такому випадку вони використовують `let`.

Однак важливо розуміти, що об’єкти (включно з масивами та функціями), призначені змінній за допомогою `const`, все ще можна мутувати. Використання оголошення `const` лише забороняє переназначення ідентифікатора змінної.

```js
const s = [5, 6, 7];
s = [1, 2, 3];
s[2] = 45;
console.log(s);
```

`s = [1, 2, 3]` призведе до помилки. Після коментування цього рядка, `console.log` відобразить значення `[5, 6, 45]`.

Як бачите, ви можете мутувати сам об’єкт `[5, 6, 7]`, і змінна `s` все одно буде вказувати на змінений масив `[5, 6, 45]`. Як і всі масиви, елементи масиву в `s` є мутованими, але через те, що було використано `const`, ви не можете використовувати ідентифікатор змінної `s`, щоб вказати на інший масив за допомогою оператора присвоєння.

# --instructions--

Масив оголошено як `const s = [5, 7, 2]`. Змініть масив на `[2, 5, 7]` за допомогою різних присвоєнь елементам.

# --hints--

Вам не слід замінювати ключове слово `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`s` має бути константною змінною (за допомогою `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+s/g));
```

Вам не слід змінювати початкове оголошення масиву.

```js
assert(__helpers.removeJSComments(code).match(
/const\s+s\s*=\s*\[\s*5\s*,\s*7\s*,\s*2\s*\]\s*;?/g
));
```

`s` має дорівнювати `[2, 5, 7]`.

```js
assert.deepEqual(s, [2, 5, 7]);
```

# --seed--

## --seed-contents--

```js
const s = [5, 7, 2];
function editInPlace() {
  // Only change code below this line

  // Using s = [2, 5, 7] would be invalid

  // Only change code above this line
}
editInPlace();
```

# --solutions--

```js
const s = [5, 7, 2];
function editInPlace() {
  s[0] = 2;
  s[1] = 5;
  s[2] = 7;
}
editInPlace();
```
