---
id: 587d7b8c367417b2b2512b56
title: Використання export для спільного використання блоку коду
challengeType: 1
forumTopicId: 301219
dashedName: use-export-to-share-a-code-block
---

# --description--

Уявіть файл під назвою `math_functions.js`, який містить кілька функцій, пов’язаних із математичними операціями. Одна з них збережена у змінній `add`, яка приймає два числа і повертає їх суму. Ви хочете використовувати цю функцію в кількох різних JavaScript файлах. Щоб поділитися нею з іншими файлами, спочатку потрібно `export` її.

```js
export const add = (x, y) => {
  return x + y;
}
```

Вищенаведений спосіб є загальним для експорту однієї функції, але те саме можна зробити так:

```js
const add = (x, y) => {
  return x + y;
}

export { add };
```

Коли ви експортуєте змінну або функцію, ви можете імпортувати її в іншому файлі і використовувати без необхідності переписувати код. Ви можете експортувати кілька речей, повторюючи перший приклад для кожної речі, яку хочете експортувати, або розмістити їх усі в інструкції export другого прикладу, ось так:

```js
export { add, subtract };
```

# --instructions--

В редакторі є дві функції, пов’язані з рядками. Експортуйте обидві, використовуючи обраний вами метод.

# --hints--

Ви повинні правильно експортувати `uppercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+uppercaseString|export\s*{\s*(uppercaseString[^}]*|[^,]*,\s*uppercaseString\s*)})/g
  )
);
```

Ви повинні правильно експортувати `lowercaseString`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /(export\s+const\s+lowercaseString|export\s*{\s*(lowercaseString[^}]*|[^,]*,\s*lowercaseString\s*)})/g
  )
);
```

# --seed--

## --seed-contents--

```js
const uppercaseString = (string) => {
  return string.toUpperCase();
}

const lowercaseString = (string) => {
  return string.toLowerCase()
}
```

# --solutions--

```js
export const uppercaseString = (string) => {
  return string.toUpperCase();
}

export const lowercaseString = (string) => {
  return string.toLowerCase()
}
```
