---
id: 598f48a36c8c40764b4e52b3
title: Запобігання мутації об’єкта
challengeType: 1
forumTopicId: 301207
dashedName: prevent-object-mutation
---

# --description--

Як було показано в попередньому завданні, одне лише оголошення `const` не захищає ваші дані від мутації. Щоб гарантувати, що ваші дані не зміняться, JavaScript надає функцію `Object.freeze` для запобігання мутації даних.

Будь-яка спроба змінити об’єкт буде відхилена, а якщо скрипт працює в суворому режимі, буде викинуто помилку.

```js
let obj = {
  name:"FreeCodeCamp",
  review:"Awesome"
};
Object.freeze(obj);
obj.review = "bad";
obj.newProp = "Test";
console.log(obj); 
```

Присвоєння `obj.review` і `obj.newProp` призведуть до помилок, оскільки наш редактор за замовчуванням працює в суворому режимі, а в консолі буде показано значення `{ name: "FreeCodeCamp", review: "Awesome" }`.

# --instructions--

У цьому завданні ви використаєте `Object.freeze` для запобігання змінам математичних констант. Вам потрібно заморозити об’єкт `MATH_CONSTANTS`, щоб ніхто не міг змінити значення `PI`, додати або видалити пропи.

# --hints--

Не слід замінювати ключове слово `const`.

```js
assert(__helpers.removeJSComments(code).match(/const/g));
```

`MATH_CONSTANTS` має бути константною змінною (за допомогою `const`).

```js
assert(__helpers.removeJSComments(code).match(/const\s+MATH_CONSTANTS/g));
```

Не слід змінювати початкове оголошення `MATH_CONSTANTS`.

```js
assert(__helpers.removeJSComments(code).match(
   /const\s+MATH_CONSTANTS\s+=\s+{\s+PI:\s+3.14\s+};/g
));
```

`PI` має дорівнювати `3.14`.

```js
assert(PI === 3.14);
```

# --seed--

## --seed-contents--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  // Only change code below this line


  // Only change code above this line
  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```

# --solutions--

```js
function freezeObj() {
  const MATH_CONSTANTS = {
    PI: 3.14
  };
  Object.freeze(MATH_CONSTANTS);

  try {
    MATH_CONSTANTS.PI = 99;
  } catch(ex) {
    console.log(ex);
  }
  return MATH_CONSTANTS.PI;
}
const PI = freezeObj();
```
