---
id: 5cdafbb0291309899753167f
title: Створіть проміс у JavaScript
challengeType: 1
forumTopicId: 301197
dashedName: create-a-javascript-promise
---

# --description--

Проміс у JavaScript — це саме те, що ви очікуєте: ви використовуєте його, щоб пообіцяти щось зробити, зазвичай асинхронно. Коли завдання завершується, ви або виконуєте свою обіцянку, або ні. `Promise` — це функція-конструктор, тому потрібно використати ключове слово `new`, щоб створити проміс. Вона приймає функцію як аргумент із двома параметрами — `resolve` і `reject`. Це методи, які визначають результат промісу. Синтаксис виглядає так:

```js
const myPromise = new Promise((resolve, reject) => {

});
```

# --instructions--

Створіть новий проміс з ім’ям `makeServerRequest`. Передайте функцію з параметрами `resolve` і `reject` у конструктор.

# --hints--

Ви повинні призначити проміс оголошеній змінній з ім’ям `makeServerRequest`.

```js
assert(makeServerRequest instanceof Promise);
```

Ваш проміс має отримувати функцію з параметрами `resolve` і `reject`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /Promise\s*\(\s*(function\s*\(\s*resolve\s*,\s*reject\s*\)\s*{|\(\s*resolve\s*,\s*reject\s*\)\s*=>\s*{)[^}]*}/g
  )
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {

});
```
