---
id: 5cdafbd72913098997531681
title: Обробка виконаного промісу за допомогою then
challengeType: 1
forumTopicId: 301203
dashedName: handle-a-fulfilled-promise-with-then
---

# --description--

Проміси найбільш корисні, коли у вашому коді є процес, який займає невідомий час (тобто щось асинхронне), часто це запит до сервера. Коли ви робите запит до сервера, це займає певний час, і після його завершення ви зазвичай хочете щось зробити з відповіддю від сервера. Це можна досягти, використовуючи метод `then`.

```js
Promise.prototype.then(onFulfilled, onRejected)
```

Метод `then` планує функції зворотного виклику для майбутнього завершення промісу — або виконання, або відхилення. Один із обробників `onFulfilled` та `onRejected` буде виконаний для обробки виконання або відхилення поточного промісу. Коли проміс виконується з `resolve`, викликається обробник `onFulfilled`.

```js
myPromise.then(result => {
  
});
```

`result` походить із аргументу, переданого методу `resolve`.

# --instructions--

Додайте метод `then` до свого промісу. Використайте `result` як параметр функції зворотного виклику та виведіть `result` у консоль.

# --hints--

Ви повинні викликати метод `then` на промісі.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/(makeServerRequest|\))\.then\(/g)
);
```

Ваш метод `then` має мати функцію зворотного виклику з параметром `result`.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  )
);
```

Ви повинні вивести `result` у консоль.

```js
assert(
  /\.then\((function\(result\){|result|\(result\)=>)/.test(
    __helpers.removeWhiteSpace(__helpers.removeJSComments(code))
  ) &&
    __helpers
      .removeWhiteSpace(__helpers.removeJSComments(code))
      .match(/\.then\(.*?result.*?console.log\(result\).*?\)/)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer is set to true to represent a successful response from a server
  let responseFromServer = true;
    
  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});

makeServerRequest.then(result => {
  console.log(result);
});
```
