---
id: 5cdafbc32913098997531680
title: Завершіть проміс за допомогою resolve і reject
challengeType: 1
forumTopicId: 301196
dashedName: complete-a-promise-with-resolve-and-reject
---

# --description--

Проміс має три стани: `pending`, `fulfilled` і `rejected`. Проміс, який ви створили в останньому завданні, назавжди застряг у стані `pending`, тому що ви не додали спосіб завершити проміс. Параметри `resolve` і `reject`, які передаються аргументу промісу, використовуються для цього. `resolve` застосовується, коли ви хочете, щоб ваш проміс успішно завершився, а `reject` — коли ви хочете, щоб він завершився з помилкою. Це методи, які приймають аргумент, як показано нижче.

```js
const myPromise = new Promise((resolve, reject) => {
  if(condition here) {
    resolve("Promise was fulfilled");
  } else {
    reject("Promise was rejected");
  }
});
```

У наведеному прикладі для аргументів цих функцій використовуються рядки, але це може бути будь-що. Часто це може бути об’єкт, дані з якого ви використовуєте для свого сайту чи іншого місця.

# --instructions--

Зробіть так, щоб проміс обробляв успіх і помилку. Якщо `responseFromServer` дорівнює `true`, викличте метод `resolve` для успішного завершення промісу. Передайте `resolve` рядок зі значенням `We got the data`. Якщо `responseFromServer` дорівнює `false`, замість цього використайте метод `reject` і передайте йому рядок: `Data not received`.

# --hints--

`resolve` слід викликати з очікуваним рядком, коли умова `if` дорівнює `true`.

```js
assert(
  __helpers.removeJSComments(code).match(/if\s*\(\s*responseFromServer\s*\)\s*{\s*resolve\s*\(\s*('|"|`)We got the data\1\s*\)(\s*|\s*;\s*)}/g)
);
```

`reject` слід викликати з очікуваним рядком, коли умова `if` дорівнює `false`.

```js
assert(
  __helpers.removeJSComments(code).match(/}\s*else\s*{\s*reject\s*\(\s*('|"|`)Data not received\1\s*\)(\s*|\s*;\s*)}/g)
);
```

# --seed--

## --seed-contents--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;
    
  if(responseFromServer) {
    // Change this line
  } else {  
    // Change this line
  }
});
```

# --solutions--

```js
const makeServerRequest = new Promise((resolve, reject) => {
  // responseFromServer represents a response from a server
  let responseFromServer;

  if(responseFromServer) {
    resolve("We got the data");
  } else {  
    reject("Data not received");
  }
});
```
