---
id: 5cddbfd622f1a59093ec611d
title: створити скрипт модуля
challengeType: 6
forumTopicId: 301198
dashedName: create-a-module-script
---

# --description--

JavaScript починав із невеликої ролі на переважно HTML-вебсайті. Сьогодні він величезний, і деякі вебсайти побудовані майже повністю на JavaScript. Щоб зробити JavaScript більш модульним, чистим і підтримуваним, ES6 представив спосіб легко ділитися кодом між JavaScript-файлами. Це передбачає експорт частин файлу для використання в одному або кількох інших файлах і імпорт потрібних частин там, де вони потрібні. Щоб скористатися цією функціональністю, потрібно створити скрипт у вашому HTML-документі з `type` типом `module`. Ось приклад:

```html
<script type="module" src="filename.js"></script>
```

Скрипт, який використовує цей тип `module`, тепер може використовувати функції `import` і `export`, про які ви дізнаєтеся в наступних завданнях.

# --instructions--

Додайте скрипт до HTML-документа з типом `module` і вкажіть для нього файл-джерело `index.js`

# --hints--

Ви повинні створити тег `script`.

```js
assert(__helpers.removeJSComments(code).match(/<\s*script[^>]*>\s*<\/\s*script\s*>/g));
```

Ваш тег `script` має мати атрибут `type` зі значенням `module`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^t]*type\s*=\s*('|")module\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

Ваш тег `script` має мати `src` зі значенням `index.js`.

```js
assert(
  __helpers.removeJSComments(code).match(
    /<\s*script\s+[^s]*src\s*=\s*('|")index\.js\1[^>]*>\s*<\/\s*script\s*>/g
  )
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <!-- Only change code below this line -->

    <!-- Only change code above this line -->
  </body>
</html>
```

# --solutions--

```html
<html>
  <body>
    <script type="module" src="index.js"></script>
  </body>
</html>
```
