---
id: 5cfa550e84205a357704ccb6
title: Використання деструктурованого присвоєння для вилучення значень з об’єктів
challengeType: 1
forumTopicId: 301216
dashedName: use-destructuring-assignment-to-extract-values-from-objects
---

# --description--

<dfn>Деструктуроване присвоєння</dfn> — це спеціальний синтаксис, введений у ES6, який дозволяє акуратно призначати значення, взяті безпосередньо з об’єкта.

Розглянемо наступний код ES5:

```js
const user = { name: 'John Doe', age: 34 };

const name = user.name;
const age = user.age;
```

`name` матиме значення рядка `John Doe`, а `age` матиме число `34`.

Ось еквівалентна інструкція присвоєння з використанням синтаксису деструктуризації ES6:

```js
const { name, age } = user;
```

Знову ж таки, `name` матиме значення рядка `John Doe`, а `age` матиме число `34`.

Тут змінні `name` і `age` будуть створені та призначені відповідними значеннями з об’єкта `user`. Ви можете побачити, наскільки це чистіше.

Ви можете вилучити стільки значень з об’єкта, скільки хочете.

# --instructions--

Замініть два присвоєння на еквівалентне деструктуроване присвоєння. Воно має призначити змінним `today` і `tomorrow` значення `today` і `tomorrow` з об’єкта `HIGH_TEMPERATURES`.

# --hints--

Вам слід видалити синтаксис присвоєння ES5.

```js
assert(
  !__helpers.removeJSComments(code).match(/today\s*=\s*HIGH_TEMPERATURES\.(today|tomorrow)/g)
);
```

Використайте деструктуризацію, щоб створити змінну `today`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(today[^}]*|[^,]*,\s*today)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

Використайте деструктуризацію, щоб створити змінну `tomorrow`.

```js
assert(
  __helpers.removeJSComments(code).match(/(var|let|const)\s*{\s*(tomorrow[^}]*|[^,]*,\s*tomorrow)\s*}\s*=\s*HIGH_TEMPERATURES(;|\s+|\/\/)/g)
);
```

`today` має бути рівним `77`, а `tomorrow` має бути рівним `80`.

```js
assert(today === 77 && tomorrow === 80);
```

# --seed--

## --seed-contents--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

// Only change code below this line

const today = HIGH_TEMPERATURES.today;
const tomorrow = HIGH_TEMPERATURES.tomorrow;

// Only change code above this line
```

# --solutions--

```js
const HIGH_TEMPERATURES = {
  yesterday: 75,
  today: 77,
  tomorrow: 80
};

const { today, tomorrow } = HIGH_TEMPERATURES;
```
