---
id: 673b567e3ba535dda140d278
title: بناء نظام تصويت
challengeType: 26
dashedName: build-a-voting-system
saveSubmissionToDB: true
---

# --description--

في هذا المختبر، ستبني نظام تصويت يستخدم `Map` لإنشاء استفتاء و`Set` لمنع التصويت المكرر.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تهيئ متغير `poll` إلى كائن جديد من نوع `Map`.

2. يجب أن تملك دالة `addOption` تقبل معلمة `option`.

3. في دالة `addOption`:

- إذا لم يكن `option` موجودًا بالفعل في الاستفتاء، يجب إضافته إلى الاستفتاء مع `Set` فارغ كقيمته لتتبع المصوتين. يجب أيضًا أن تُرجع الرسالة `Option "<option>" added to the poll.`

- إذا كان `option` موجودًا بالفعل، يجب أن تُرجع الرسالة `Option "<option>" already exists.`

- إذا حاولت إضافة خيار فارغ، يجب أن تُرجع الدالة الرسالة `Option cannot be empty.`

4. يجب أن تملك دالة `vote` تقبل معلمتين، `option` (الخيار للتصويت له) و`voterId` (معرف فريد للمصوت).

5. في دالة `vote`:

- إذا لم يكن `option` موجودًا في الاستفتاء، يجب أن تُرجع الدالة الرسالة `Option "<option>" does not exist.`

- إذا كان `option` موجودًا، يجب على الدالة التحقق مما إذا كان `voterId` قد صوت بالفعل لهذا `option`.

- إذا كان المصوت قد صوت بالفعل، يجب أن تُرجع الدالة الرسالة `Voter <voterId> has already voted for "<option>".`

- إذا لم يصوّت الناخب من قبل، يجب إضافة `voterId` إلى `Set` الخاص بـ`voters` لهذه الخيار. يجب أن تُرجع الدالة الرسالة `Voter <voterId> voted for "<option>".`

6. يجب أن يكون لديك على الأقل ثلاثة خيارات في `poll`.

7. يجب أن يحتوي `poll` على ثلاثة أصوات على الأقل.

8. يجب أن تملك دالة `displayResults` تُرجع نتائج الاستفتاء بالتنسيق التالي:

```js
Poll Results:
OptionA: N votes
OptionB: N votes
.
.

/*
sample output

Poll Results:
Turkey: 2 votes
Morocco: 1 votes
*/
```

# --hints--

يجب أن تهيئ متغير `poll` إلى كائن جديد من نوع `Map`.

```js
assert.exists(poll);
assert.instanceOf(poll, Map);
```

يجب أن تعرف دالة `addOption` تقبل معلمة `option`.

```js
assert.isFunction(addOption);
assert.match(
  addOption.toString(),
  /^function\s+addOption\s*\(\s*option\s*\)\s*\{|\(?\s*option\s*\)?\s*=>/
);
```

يجب أن تعرف دالة `vote` تقبل معلمتين، `option` و`voterId`.

```js
assert.isFunction(vote);
assert.match(vote.toString(), /\(\s*option\s*,\s*voterId\s*\)/);
```

يجب أن تعرف دالة `displayResults` لعرض نتائج الاستفتاء.

```js
assert.isFunction(displayResults);
```

يجب أن يكون لديك على الأقل ثلاثة خيارات في `poll`.

```js
assert.isAtLeast(poll.size, 3);
```

يجب أن يحتوي `poll` على ثلاثة أصوات على الأقل.

```js
let totalVotes = 0;
for (let voters of poll.values()) {
  totalVotes += voters.size;
}
assert.isAtLeast(totalVotes, 3);
```

يجب أن تضمن أن كل خيار تصويت مرتبط بكائن `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();

  addOption('Turkey');
  assert.instanceOf(poll.get('Turkey'), Set);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`addOption("Egypt")` يجب أن تُرجع `Option "Egypt" added to the poll.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(addOption('Egypt'), 'Option "Egypt" added to the poll.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

إضافة خيار فارغ يجب أن تُرجع `"Option cannot be empty."`

```js
assert.equal(addOption(''), 'Option cannot be empty.');
```

عندما يكون `Turkey` مضافًا بالفعل، يجب أن تُرجع `addOption("Turkey")` الرسالة `Option "Turkey" already exists.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Turkey');
  assert.equal(addOption('Turkey'), 'Option "Turkey" already exists.');
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

عندما يكون `Malaysia` موجودًا في خيارات التصويت، يجب أن تُرجع `vote("Malaysia", "traveler1")` الرسالة `Voter traveler1 voted for "Malaysia".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Malaysia');
  assert.equal(
    vote('Malaysia', 'traveler1'),
    'Voter traveler1 voted for "Malaysia".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`vote` يجب أن تُحدث `Set` الخاص بالمصوتين لخيار ما.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Alaska');
  vote('Alaska', 'seal');
  assert.isTrue(poll.get('Alaska').has('seal'));
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

عندما يحاول `traveler1` التصويت لـ `Algeria` مرة أخرى، يجب أن تُرجع `vote("Algeria", "traveler1")` الرسالة `Voter traveler1 has already voted for "Algeria".`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Algeria');
  vote('Algeria', 'traveler1');
  assert.equal(
    vote('Algeria', 'traveler1'),
    'Voter traveler1 has already voted for "Algeria".'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

الأصوات المكررة لا يجب أن تزيد حجم `Set`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  vote('Turkey', 'traveler1');
  const voterSet = poll.get('Turkey');
  const initialSize = voterSet.size;

  vote('Turkey', 'traveler1');
  assert.equal(voterSet.size, initialSize);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

عندما لا يكون `Nigeria` من ضمن خيارات التصويت، يجب أن تُرجع `vote("Nigeria", "traveler2")` الرسالة `Option "Nigeria" does not exist.`

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  assert.equal(
    vote('Nigeria', 'traveler2'),
    'Option "Nigeria" does not exist.'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

يجب أن يكون الخيار الفريد قادرًا على استقبال عدة أصوات.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Bali');
  assert.equal(vote('Bali', 'traveler1'), 'Voter traveler1 voted for "Bali".');
  assert.equal(vote('Bali', 'traveler2'), 'Voter traveler2 voted for "Bali".');
  assert.equal(poll.get('Bali').size, 2);
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

عندما يحصل خيار `Turkey` على صوتين، وخيار `Morocco` على صوت واحد، وخيار `Spain` لا يحصل على أي صوت، يجب أن تُرجع `displayResults()` الرسالة `"Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes"`.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');
  addOption('Spain');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes\nSpain: 0 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

`displayResults()` يجب أن تُرجع النتائج بالتنسيق الصحيح.

```js
const pollCopy = new Map(poll);
try {
  poll.clear();
  addOption('Turkey');
  addOption('Morocco');

  vote('Turkey', 'traveler1');
  vote('Turkey', 'traveler2');
  vote('Morocco', 'traveler3');

  assert.equal(
    displayResults(),
    'Poll Results:\nTurkey: 2 votes\nMorocco: 1 votes'
  );
} finally {
  poll.clear();
  pollCopy.forEach((val, key) => {
    poll.set(key, val);
  });
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const poll = new Map();

function addOption(option) {
  if (!option || option.trim() === '') {
    return 'Option cannot be empty.';
  }
  if (!poll.has(option)) {
    poll.set(option, new Set());
    return `Option "${option}" added to the poll.`;
  } else {
    return `Option "${option}" already exists.`;
  }
}

function vote(option, voterId) {
  if (!poll.has(option)) {
    return `Option "${option}" does not exist.`;
  }
  const voters = poll.get(option);
  if (voters.has(voterId)) {
    return `Voter ${voterId} has already voted for "${option}".`;
  } else {
    voters.add(voterId);
    return `Voter ${voterId} voted for "${option}".`;
  }
}

function displayResults() {
  let results = 'Poll Results:\n';
  for (let [option, voters] of poll.entries()) {
    results += `${option}: ${voters.size} votes\n`;
  }
  return results.trim();
}

addOption('Turkey');
addOption('Morocco');
addOption('Spain');

vote('Turkey', 'traveler1');
vote('Turkey', 'traveler2');
vote('Morocco', 'traveler3');
```
