---
id: 67051431a73f1ca25d9a6f25
title: 프로젝트 아이디어 보드 만들기
challengeType: 26
dashedName: build-a-project-idea-board
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `projectStatus`라는 이름의 객체 상수를 세 개의 키 `PENDING`, `SUCCESS`, `FAILURE`로 정의하세요. 각 상태에는 `description` 키가 있는 객체가 할당되어야 하며, 값은 각각 `Pending Execution`, `Executed Successfully`, `Execution Failed`입니다.
2. `ProjectIdea`라는 이름의 클래스를 정의하세요. 이 클래스에는 `constructor`과 `title`을 매개변수로 받는 `description`가 있어야 합니다. 제공된 매개변수로 `title`와 `description` 속성을 초기화하세요. 또한 기본적으로 `status` 값으로 설정된 `projectStatus.PENDING`라는 속성이 있어야 합니다.
3. `updateProjectStatus` 클래스 안에 `ProjectIdea`라는 이름의 메서드를 정의하세요. 이 메서드는 `newStatus` 매개변수를 받아 `status` 속성을 해당 값으로 업데이트해야 합니다.
4. `ProjectIdeaBoard` 클래스를 정의하세요. 이 클래스에는 `constructor`를 매개변수로 받는 `title`가 있어야 하며, `ideas` 클래스 인스턴스를 담을 빈 배열 `ProjectIdea`을 초기화해야 합니다.
5. `pin` 클래스 안에 `ProjectIdeaBoard`라는 이름의 메서드를 정의하세요. 이 메서드는 `ProjectIdea` 클래스의 인스턴스를 받아 해당 인스턴스를 `ideas` 배열에 추가해야 합니다.
6. `unpin` 클래스 안에 `ProjectIdeaBoard`라는 이름의 메서드를 정의하세요. 이 메서드는 `ProjectIdea` 클래스의 인스턴스를 받아 `ideas` 배열에서 제거해야 합니다.
7. `count` 배열에 있는 프로젝트 아이디어 수를 반환하는 `ProjectIdeaBoard`라는 이름의 메서드를 정의하세요.
8. 프로젝트 아이디어의 이름, 설명, 상태를 다음 형식으로 반환하는 `formatToString`라는 이름의 메서드를 정의하세요:

```js
<ProjectIdeaBoard title> has <ProjectIdeaBoard count> idea(s)
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
<ProjectIdea title> (<ProjectIdea status description>) - <ProjectIdea description>
.
.
.
```

# --hints--

올바른 프로젝트 상태와 설명을 가진 `projectStatus`라는 이름의 객체 상수를 정의하세요.

```js
assert.isObject(projectStatus);
assert.sameMembers(Object.keys(projectStatus), ["PENDING", "SUCCESS", "FAILURE"]);
assert.deepEqual(projectStatus.PENDING, { description: "Pending Execution" });
assert.deepEqual(projectStatus.SUCCESS, { description: "Executed Successfully" });
assert.deepEqual(projectStatus.FAILURE, { description: "Execution Failed" });
```

`ProjectIdea` 클래스를 정의하세요.

```js
assert.isFunction(ProjectIdea);
assert.isObject(new ProjectIdea());
```

`ProjectIdea` 클래스는 전달된 매개변수에 따라 `title` 속성과 `description` 속성을 초기화해야 합니다. 또한 상태를 `projectStatus.PENDING`로 설정해야 합니다.

```js
const test = new ProjectIdea("sample title", "sample description");
assert.isAbove(test.title.length, 0);
assert.isAbove(test.description.length, 0);
assert.strictEqual(test.status.description, "Pending Execution");
```

`ProjectIdea` 클래스에는 `updateProjectStatus` 메서드가 있어야 합니다.

```js
const projectIdea = new ProjectIdea();
assert.isFunction(projectIdea.updateProjectStatus);
```

`ProjectIdeaBoard` 클래스를 정의하세요.

```js
assert.isFunction(ProjectIdeaBoard);
assert.isObject(new ProjectIdeaBoard());
```

`ProjectIdeaBoard`는 전달된 매개변수에 따라 `title` 속성을 초기화하고, `ideas` 클래스 인스턴스를 담을 빈 배열 `ProjectIdea`을 초기화해야 합니다.

```js
const test = new ProjectIdeaBoard("Sample Title");
assert.isArray(test.ideas);
assert.lengthOf(test.ideas, 0);
assert.exists(test.title);
```

`ProjectIdeaBoard` 클래스에는 `pin` 메서드가 있어야 합니다.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.pin);
```

`ProjectIdeaBoard` 클래스에는 `unpin` 메서드가 있어야 합니다.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.unpin);
```

`ProjectIdeaBoard` 클래스에는 `count` 메서드가 있어야 합니다.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.count);
```

`ProjectIdeaBoard` 클래스에는 `formatToString` 메서드가 있어야 합니다.

```js
const projectIdeaBoard = new ProjectIdeaBoard();
assert.isFunction(projectIdeaBoard.formatToString);
```

`new ProjectIdea("Smart Window Locks", "An automation project allowing users to lock, unlock windows automatically based on weather conditions.")`는 `{ title: 'Smart Window Locks', description: 'An automation project allowing users to lock, unlock windows automatically based on weather conditions.', status: { description: 'Pending Execution' } }`를 반환해야 합니다.

```js
const windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);
assert.deepEqual(windowLocks, {
  title: "Smart Window Locks",
  description:
    "An automation project allowing users to lock, unlock windows automatically based on weather conditions.",
  status: { description: "Pending Execution" },
});
```

`updateProjectStatus(projectStatus.SUCCESS)`에서 `new ProjectIdea("Fitness Tracker App", "An app that tracks user workouts, diet, and sleep patterns.")`를 호출하면 상태가 `{ description: 'Executed Successfully' }`로 업데이트되어야 합니다.

```js
let mobileApp = new ProjectIdea(
  "Fitness Tracker App",
  "An app that tracks user workouts, diet, and sleep patterns."
);

// Update the project status to SUCCESS
mobileApp.updateProjectStatus(projectStatus.SUCCESS);
assert.deepEqual(mobileApp, {
  title: "Fitness Tracker App",
  description: "An app that tracks user workouts, diet, and sleep patterns.",
  status: { description: "Executed Successfully" },
});
```

`updateProjectStatus(projectStatus.FAILURE)`에서 `new ProjectIdea("Breakfast Chef Robot", "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.")`를 호출하면 상태가 `{ description: 'Execution Failed' }`로 업데이트되어야 합니다.

```js
let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);

// Update the project status to FAILURE
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

assert.deepEqual(breakfastChef, {
  title: "Breakfast Chef Robot",
  description:
    "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone.",
  status: { description: "Execution Failed" },
});
```

`updateProjectStatus(projectStatus.SUCCESS)`에서 `new ProjectIdea("Online Used Video Games Store", "An online platform where users can buy second hand physical copies of video games from other users.")`를 호출하면 상태가 `{ description: 'Executed Successfully' }`로 업데이트되어야 합니다.

```js
let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);

// Update the project status set to SUCCESS
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

assert.deepEqual(usedVideoGamesStore, {
  title: "Online Used Video Games Store",
  description:
    "An online platform where users can buy second hand physical copies of video games from other users.",
  status: { description: "Executed Successfully" },
});
```

`ProjectIdea`에 `ProjectIdeaBoard` 객체를 `pin` 메서드를 사용해 고정할 수 있어야 합니다.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin the idea
let techProjects = new ProjectIdeaBoard("Tech Projects Board");
techProjects.pin(smartHome);

// Correct the assertion to compare the first idea in the array
assert.deepEqual(techProjects.ideas[0], {
    title: 'Smart Home System',
    description: 'An integrated system to control lighting, temperature, and security devices remotely.',
    status: { description: 'Pending Execution' }
});
```

`ProjectIdea`에서 `ProjectIdeaBoard` 객체를 `unpin` 메서드를 사용해 고정 해제할 수 있어야 합니다.

```js
let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

techProjects.pin(smartHome);
techProjects.unpin(smartHome);

assert.isEmpty(techProjects.ideas);
```

`new ProjectIdeaBoard("Empty Board")`가 비어 있을 때 `emptyBoard.formatToString()`는 `Empty Board has 0 idea(s)\n`를 반환해야 합니다.

```js

let smartHome = new ProjectIdea(
  "Smart Home System",
  "An integrated system to control lighting, temperature, and security devices remotely."
);

// Create a project board and pin ideas
let techProjects = new ProjectIdeaBoard("Tech Projects Board");

// Create an empty board
let emptyBoard = new ProjectIdeaBoard("Empty Board");

// Attempt to unpin a project from an empty board
emptyBoard.unpin(smartHome);

assert.equal(emptyBoard.formatToString(), "Empty Board has 0 idea(s)\n");
```

`new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.")`에 `new ProjectIdeaBoard("Tech Projects Board")`를 고정하면 `techProjects.formatToString()`는 `Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n`를 반환해야 합니다.

```js
let techProjectsBoard = new ProjectIdeaBoard("Tech Projects Board");
let smartHome = new ProjectIdea("Smart Home System", "An integrated system to control lighting, temperature, and security devices remotely.");
techProjectsBoard.pin(smartHome);

assert.equal(techProjectsBoard.formatToString(), "Tech Projects Board has 1 idea(s)\nSmart Home System (Pending Execution) - An integrated system to control lighting, temperature, and security devices remotely.\n"); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const projectStatus = {
  PENDING: { description: "Pending Execution" },
  SUCCESS: { description: "Executed Successfully" },
  FAILURE: { description: "Execution Failed" },
};

class ProjectIdea {
  constructor(title, description) {
    this.title = title;
    this.description = description;
    this.status = projectStatus.PENDING;
  }

  updateProjectStatus(newStatus) {
    this.status = newStatus;
  }
}

class ProjectIdeaBoard {
  constructor(title) {
    this.title = title;
    this.ideas = [];
  }

  pin(newIdea) {
    this.ideas.push(newIdea);
  }

  unpin(idea) {
    this.ideas.splice(this.ideas.indexOf(idea), 1);
  }

  count() {
    return this.ideas.length;
  }

  formatToString() {
    let representation = `${this.title} has ${this.count()} idea(s)\n`;

    this.ideas.forEach(
      (idea) =>
        (representation += `${idea.title} (${idea.status.description}) - ${idea.description}\n`),
    );
    return representation;
  }
}

let windowLocks = new ProjectIdea(
  "Smart Window Locks",
  "An automation project allowing users to lock, unlock windows automatically based on weather conditions."
);

let breakfastChef = new ProjectIdea(
  "Breakfast Chef Robot",
  "A robot that can follow a given list of instructions and prepare breakfast for the user and let them know through their phone."
);
breakfastChef.updateProjectStatus(projectStatus.FAILURE);

let usedVideoGamesStore = new ProjectIdea(
  "Online Used Video Games Store",
  "An online platform where users can buy second hand physical copies of video games from other users."
);
usedVideoGamesStore.updateProjectStatus(projectStatus.SUCCESS);

let hardwareProjectIdeas = new ProjectIdeaBoard("Hardware Project Board");
hardwareProjectIdeas.pin(windowLocks);
hardwareProjectIdeas.pin(breakfastChef);

let softwareProjectIdeas = new ProjectIdeaBoard("Software Project Board");
softwareProjectIdeas.pin(usedVideoGamesStore);

console.log(softwareProjectIdeas.formatToString());

console.log(hardwareProjectIdeas.formatToString());
hardwareProjectIdeas.unpin(breakfastChef);
console.log(hardwareProjectIdeas.formatToString());
```
