---
id: 657e928716b77b2277980276
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

在这个项目中，你将构建一个程序，该程序以 `camelCase` 或 `PascalCase` 格式的字符串作为输入，并使用两种方法将其转换为 `snake_case` 格式的字符串。 首先，你将使用 for 循环，然后使用列表理解来实现相同的结果。 你将看到列表推导如何使你的代码更加简洁。

开始定义一个名为 `convert_to_snake_case()` 的新函数，该函数接受名为 `pascal_or_camel_cased_string` 的字符串作为输入。 现在，在函数内添加一个 `pass` 语句。

# --hints--

你应该定义一个名为 `convert_to_snake_case()` 的新函数。 不要忘记最后的冒号，并使用 `pass` 来填充函数体。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(convert_to_snake_case)
    `));
  }
})
```

你的函数应该采用一个名为 `pascal_or_camel_cased_string` 的参数。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      'pascal_or_camel_cased_string' in inspect.signature(convert_to_snake_case).parameters
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
