---
id: 657ee28cefc4945568287673
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

有了空列表后，现在你可以开始遍历输入字符串并将其转换为蛇形大小写。

在函数内部，在刚刚创建的列表下方，添加一个 `for` 循环来遍历 `pascal_or_camel_cased_string`。 确保将目标变量命名为 `char`。 现在，在循环主体中添加 `pass` 语句作为占位符。

# --hints--

你应该编写一个新的 `for` 循环，其目标变量名为 `char`。 不要忘记最后的冒号，并使用 `pass` 来填充循环体。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +for\s+char\s+in\s+pascal_or_camel_cased_string\s*:\s*pass[\s]*$/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
--fcc-editable-region--
```
