---
id: 657efa642593c5746acc5c81
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

在 `if` 语句主体中，你需要将任何大写字符转换为小写，并在该小写字符前面添加下划线。

使用 `.lower()` 字符串方法将大写字符转换为小写字符。 然后，在字符前添加下划线。 将结果分配给名为 `converted_character` 的变量。

# --hints--

你应该将修改后的字符分配给名为 `converted_character` 的变量。

```js
({
    test: () => {
   assert(runPython(`
      converted_character_variable = (
          _Node(_code)
          .find_function("convert_to_snake_case")
          .find_for_loops()[0]
          .find_ifs()[0]
          .find_variable("converted_character")
      )
      
      converted_character_variable.is_equivalent(
          "converted_character = '_' + char.lower()"
      ) or converted_character_variable.is_equivalent(
          "converted_character = f'_{char.lower()}'"
      )
    `)); 
    }
})
```

你的 `if` 语句中不应包含 `pass`。

```js
({
    test: () => {
        assert(runPython(`not _Node(_code).find_function("convert_to_snake_case").find_for_loops()[0].find_ifs()[0].has_stmt("pass")`)); 
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
--fcc-editable-region--
        if char.isupper():
            pass
--fcc-editable-region--
```
