---
id: 657ef2a86d4e545cec9a85fb
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

在驼峰式命名法和帕斯卡式命名法中，大写字母标记着新单词的开始。 要将输入字符串转换为蛇形大小写，你需要检查输入字符串中的字符是否为大写。 

你可以使用 `.isupper()` 字符串方法来检查字符是否为大写。 如果字符是大写，则此方法返回 `True`，如果不是，则返回 `False`。

在 `for` 循环内，添加 `if` 语句来检查当前字符是否为大写。 将 `pass` 语句移至新的 `if` 语句内。

# --hints--

你应该编写一个新的 if 语句，以 char.isupper() 作为条件。 记得在最后添加冒号，并使用 pass 关键字来填充 if 语句主体。

```js
({
    test: () => {        
        const commentless_code = __helpers.python.removeComments(code)
        const transformedCode = commentless_code.replace(/\r/g, "");
        const {function_body} = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        assert.match(function_body, /\s+if\s+char\.isupper\s*\(\s*\)\s*:\s+pass/);
    }
})
```

你应该将循环体中的 `pass` 语句替换为 `if` 语句。

```js
({
   test: () => {
        const commentless_code = __helpers.python.removeComments(code)
        const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+pascal_or_camel_cased_string\s*/);
        assert.notMatch(block_body, /^\s+pass/);

    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        pass
--fcc-editable-region--        
```
