---
id: 657effaa2a5e0277d71f9cbe
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

此时，变量 `snake_cased_char_list` 保存了转换后的字符列表。 要将这些字符组合成一个字符串，你可以使用 `.join()` 方法。

`join` 方法的工作原理是将列表的每个元素连接成一个字符串，并用指定的字符串（称为分隔符）分隔。

```py
result_string = ''.join(characters)
```

上面的例子将 `characters` 列表中的元素连接成一个字符串，其中每个元素使用空字符串作为分隔符连接在一起。

现在，在 `for` 循环之后，使用 `.join()` 方法以空字符串作为分隔符连接 `snake_cased_char_list` 中的元素。 将结果分配给名为 `snake_cased_string` 的新变量。


# --hints--

你应该使用 `.join()` 方法将 `snake_cased_char_list` 列表中的字符连接成一个字符串，并在 `for` 循环之后将该字符串分配给名为 `snake_cased_string` 的变量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +snake_cased_string\s*=\s*('|")\1\.join\s*\(\s*snake_cased_char_list\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--


--fcc-editable-region--
```
