---
id: 657ed53c19461d4b95c4757a
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

你需要添加一个空列表，用于保存将字符串的字符转换为蛇形命名法之后的字符。 

在函数内部，通过创建一个名为 `snake_cased_char_list` 的空列表替换 `pass` 语句。

# --hints--

你应该将 `pass` 替换为 `convert_to_snake_case` 中名为 `snake_cased_char_list` 的空列表。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const commentless_code = __helpers.python.removeComments(transformedCode)
        const convert_to_snake_case = __helpers.python.getDef("\n" + commentless_code, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;
        assert.notMatch(function_body, /\bpass\b/);
        assert.match(function_body, / +snake_cased_char_list\s*=\s*\[\s*\]\s*/);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    pass    
--fcc-editable-region--
```
