---
id: 657f456223b8c1187b461987
title: 步骤 19
challengeType: 20
dashedName: step-19
---

# --description--

在 Python 中，<dfn>列表推导式</dfn> 是一种构造，它允许你通过将表达式应用于现有可迭代对象中的每个项目并可选地使用条件过滤项目来生成新列表。 除了更简洁之外，列表推导通常运行速度更快。

基本列表推导由一个表达式和后跟 `for` 子句组成：

```py
spam = [i * 2 for i in iterable]
```

上面使用变量 `i` 来迭代 `iterable`。 通过在当前迭代中评估表达式 `i * 2` 来获得结果列表的每个元素。

在此步骤中，你需要使用列表推导语法填充空列表 `snake_cased_char_list`。

将空列表转换为列表推导，将 `pascal_or_camel_cased_string` 中的每个字符转换为小写字符，并在其前面添加下划线（之前注释掉的代码可能会帮助你编写表达式）。 使用 `char` 来迭代 `pascal_or_camel_cased_string`。 

# --hints--

你应该将 `snake_cased_char_list` 转换为对 `pascal_or_camel_cased_string` 进行迭代的列表推导。

```js
({ 
    test: () => assert(runPython(`
    iters = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_iters()
    len(iters) == 1 and iters[0].is_equivalent("pascal_or_camel_cased_string")
    `)) 
})
```

你的列表推导应该使用 `char` 来迭代 `pascal_or_camel_cased_string`。

```js
({ 
    test: () => assert(runPython(`
    targets = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_targets()
    len(targets) == 1 and targets[0].is_equivalent("char")
    `)) 
})
```

你的列表推导应该对 `pascal_or_camel_cased_string` 中的每个 `char` 评估 `'_' + char.lower()`。

```js
({ 
    test: () => assert(runPython(`
    target_var = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list")
    target_var.find_comp_expr().is_equivalent("'_' + char.lower()") and len(target_var.find_comp_ifs()) == 0
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
