---
id: 657f465f8e718b19c5105ae5
title: 步骤 20
challengeType: 20
dashedName: step-20
---

# --description--

列表推导接受条件语句，仅当满足某些条件时才评估提供的表达式：

```py
spam = [i * 2 for i in iterable if i > 0]
```

从输出中可以看到，从 `pascal_or_camel_cased_string` 生成的字符列表已经加入。 由于列表推导中的表达式针对每个字符进行评估，因此结果是一个小写字符串，所有字符都用下划线分隔。

按照上面的示例，在列表推导中添加 `if` 子句，以便仅当字符为大写时才执行表达式。

# --hints--

你应该在列表推导中添加一个带有条件 `char.isupper()` 的 `if` 子句。

```js
({ 
    test: () => assert(runPython(`
    ifs = _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_ifs()
    len(ifs) == 1 and ifs[0].is_equivalent("char.isupper()")
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
