---
id: 657f0044be09db790b1eb1c5
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

在帕斯卡格式中，字符串以大写字母开头。 将所有字符转换为小写并添加下划线后，字符串开头可能会出现额外的下划线。

解决此问题的最简单方法是使用 `.strip()` 字符串方法，该方法从字符串中删除作为其参数传递的一组字符中的任何前导或尾随字符。 例如：

```py
original_string = "_example_string_"

clean_string = original_string.strip('_')
```

`strip()` 方法应用于 `original_string`。 这将删除所有前导和尾随下划线。 上述示例的结果将是字符串 `'example_string'`。

声明一个名为 `clean_snake_cased_string` 的新变量，并将应用于 `snake_cased_string` 的 `.strip()` 方法的结果分配给它，同时将 `'_'` 作为参数传递给该方法。

# --hints--

你应该将 `snake_cased_string.strip('_')` 的结果字符串分配给名为 `clean_snake_cased_string` 的变量。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +clean_snake_cased_string\s*=\s*snake_cased_string.strip\(\s*("|')_\1\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
--fcc-editable-region--
    snake_cased_string = ''.join(snake_cased_char_list)
--fcc-editable-region--
```
