---
id: 657f0353c9523d7d896873ea
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

在 `main()` 函数内部，将 `pass` 语句替换为对 `convert_to_snake_case()` 函数的调用，并将字符串 `'aLongAndComplexString'` 作为输入传递。

要显示输出，请将函数调用作为参数传递给 `print()` 函数。

# --hints--

你应该在 `main()` 函数内调用 `convert_to_snake_case()` 并将 `'aLongAndComplexString'` 作为输入传递给该函数。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\s*\(\s*convert_to_snake_case\s*\(\s*("|')aLongAndComplexString\1\s*\)\s*\)/);
    }
})
```

你的 `main` 函数中不应有 `pass`。

```js
({
    test: () => {
        const commentless_code = __helpers.python.removeComments(code);
        const transformedCode = commentless_code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.notMatch(function_body, /pass/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string

--fcc-editable-region--
def main():
    pass
--fcc-editable-region--
```
