---
id: 657f47b12c51e41b3149e584
title: 步骤21
challengeType: 20
dashedName: step-21
---

# --description--

但最终结果并不完全是你想要达到的。 你需要对 `if` 子句过滤掉的字符执行不同的表达式。 你将使用 `else` 子句来实现这一点：

```py
spam = [i * 2 if i > 0 else -1 for i in iterable]
```

请注意，与 `if` 子句不同，`if`/`else` 结构必须放在表达式和 `for` 关键字之间。

修改列表推导，以便当字符不是大写时它保持不变。

# --hints--

你应该修改列表推导，以在 `char.isupper()` 时评估表达式 `'_' + char.lower()`，否则评估表达式 char。

```js
({ 
    test: () => assert(runPython(`
    _Node(_code).find_function("convert_to_snake_case").find_variable("snake_cased_char_list").find_comp_expr().is_equivalent("'_' + char.lower() if char.isupper() else char")
    `)) 
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

--fcc-editable-region--
    snake_cased_char_list = ['_' + char.lower() for char in pascal_or_camel_cased_string if char.isupper()]
--fcc-editable-region--
    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
