---
id: 657f425dbab54e11993c80f0
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

到目前为止，在这个项目中，你已经使用 `for` 循环来迭代输入字符串并将其转换为所需的输出。 现在你将开始从 `for` 循环转换到列表推导。

首先注释掉 `convert_to_snake_case()` 函数内的所有代码行。 不要删除它们，因为当你使用列表推导实现逻辑时它们会很有帮助。

记得在函数体中添加 `pass` 关键字，以防止代码在测试期间失败。

# --hints--

你应该注释掉 `convert_to_snake_case()` 函数内的所有 10 行代码，不要注释掉其他任何内容。 添加 `pass` 临时填充函数体，避免错误。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "").replace(/\n#/g, "\n #");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.equal(function_body.split("#").length - 1, 10);
    }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()

```
