---
id: 657f025ec86c3d7c4177b6be
title: 步骤11
challengeType: 20
dashedName: step-11
---

# --description--

函数完成后，你现在可以在另一个函数中使用它。

创建一个名为 `main()` 的新函数，并以 `pass` 作为函数主体。

# --hints--

你应该定义一个名为 `main()` 的新函数。 不要忘记最后的冒号，并使用 `pass` 来填充函数体。

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isfunction(main)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')

    return clean_snake_cased_string
--fcc-editable-region--


--fcc-editable-region--
```
