---
id: 657f01ae9aea647b27402d3e
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

要完成该函数，请返回 `clean_snake_cased_string`。 这将完成该功能并允许你使用它将字符串从帕斯卡或驼峰式大小写转换为蛇形大小写。

在函数末尾添加 `return` 语句以返回 `clean_snake_cased_string`。


# --hints--

你应该在 `convert_to_snake_case` 函数的末尾添加 `return clean_snake_cased_string`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
        const { function_body } = convert_to_snake_case;

        assert.match(function_body, / +return\s+clean_snake_cased_string/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    snake_cased_char_list = []
    for char in pascal_or_camel_cased_string:
        if char.isupper():
            converted_character = '_' + char.lower()
            snake_cased_char_list.append(converted_character)
        else:
            snake_cased_char_list.append(char)
    snake_cased_string = ''.join(snake_cased_char_list)
    clean_snake_cased_string = snake_cased_string.strip('_')
--fcc-editable-region--


--fcc-editable-region--
```
