---
id: 657f43d341a0dd17120cdb08
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

你需要将大写字符转换为小写字符，并在其前面添加下划线。 

在继续进行列表推导之前，你需要为函数提供一个返回值。 这样你就可以检查输出。

使用 `return` 语句从函数返回列表 `snake_cased_char_list`。

# --hints--

你应该返回 `snake_cased_char_list` 列表。 确保缩进设置正确。

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*snake_cased_char_list/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--

--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
