---
id: 657f4add33ea4b1f61ba3dc8
title: 步骤-23
challengeType: 20
dashedName: step-23
---

# --description--

最后通过执行程序来尝试这个新的实现。 将输入字符串更改为 `'IAmAPascalCasedString'` 并查看它是否显示为 `'i_am_a_pascal_cased_string'`，即使这是一个谎言。

如果你已正确完成所有操作，你应该会看到输入字符串转换为蛇形大小写，就像以前一样。

恭喜！ 现在你的 `convert_to_snake_case()` 函数已准备就绪。

# --hints--

你应该在 `main()` 函数中将输入字符串从 `'aLongAndComplexString'` 更改为 `'IAmAPascalCasedString'`。

```js
({
    test: () => {
        const transformedCode = code.replace(/\r/g, "");
        const main = __helpers.python.getDef("\n" + transformedCode, "main");
        const { function_body } = main;

        assert.match(function_body, / +print\(\s*convert_to_snake_case\(\s*'IAmAPascalCasedString'\s*\)\s*\)/);
    }
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

--fcc-editable-region--
def main():
    print(convert_to_snake_case('aLongAndComplexString'))

    
--fcc-editable-region--

main()
```

# --solutions--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):

    snake_cased_char_list = [
        '_' + char.lower() if char.isupper()
        else char
        for char in pascal_or_camel_cased_string
    ]

    return ''.join(snake_cased_char_list).strip('_')

def main():
    print(convert_to_snake_case('IAmAPascalCasedString'))

main()
```
