---
id: 663b10c10a4c0a0e095137ee
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

你不需要返回列表 `snake_cased_char_list`，而是需要使用空字符串 `''` 作为分隔符将其元素连接成单个字符串。

修改 `return` 语句，返回以空字符串作为分隔符，连接 `snake_cased_char_list` 的结果。

# --hints--

你应该修改 `return` 语句以返回使用 `join()` 方法连接 `snake_cased_char_list` 和空字符串 `''` 作为分隔符的结果。

```js
({
    test: () => assert(runPython(`
    _Node(_code).find_function('convert_to_snake_case').find_return().is_equivalent('return "".join(snake_cased_char_list)')
    `))
})
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return snake_cased_char_list
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
