---
id: 663b16e62fee463b4caf46e9
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

在加入列表 `snake_cased_char_list` 的元素后，你需要从结果字符串中删除任何前导或尾随下划线。 为此，使用 `strip` 方法并以下划线字符 `_` 作为参数。

方法调用可以链接在一起，这意味着一个方法调用的结果可以用作另一个方法调用的对象。

```py
words_list = ['hello', 'world', 'this', 'is', 'chained', 'methods']
result = ' '.join(words_list).upper()
```

在上面的例子中，`.upper()` 方法链接到 `' '.join(words_list)`，因此在 `.join()` 调用的结果上调用 `.upper()`。

通过将对 `.strip()` 方法的调用链接到 `''.join(snake_cased_char_list)` 来修改 `return` 语句，以删除任何前导或尾随下划线。

# --hints--

你应该修改返回语句以包含 `strip` 方法并以下划线字符 `_` 作为参数。

```js
const transformedCode = code.replace(/\r/g, "");
const convert_to_snake_case = __helpers.python.getDef("\n" + transformedCode, "convert_to_snake_case");
const { function_body } = convert_to_snake_case;

assert.match(function_body, /return\s*''.join\(snake_cased_char_list\).strip\('_'\)/);
```

# --seed--

## --seed-contents--

```py
def convert_to_snake_case(pascal_or_camel_cased_string):
    # snake_cased_char_list = []
    # for char in pascal_or_camel_cased_string:
    #     if char.isupper():
    #       converted_character = '_' + char.lower()
    #       snake_cased_char_list.append(converted_character)
    #     else:
    #         snake_cased_char_list.append(char)
    # snake_cased_string = ''.join(snake_cased_char_list)
    # clean_snake_cased_string = snake_cased_string.strip('_')

    # return clean_snake_cased_string

    snake_cased_char_list = []
--fcc-editable-region--
    return ''.join(snake_cased_char_list)
--fcc-editable-region--

def main():
    print(convert_to_snake_case('aLongAndComplexString'))

main()
```
