---
id: 614387cbefeeba5f3654a291
title: 步骤6
challengeType: 0
dashedName: step-6
---

# --description--

在 `header` 元素之后，创建一个 `div`，并将 `class` 设置为 `author`。

在这个 `div` 中，创建一个 `p` 元素，将 `class` 设置为 `author-name`，并给它添加文本 `By freeCodeCamp`。 将 `freeCodeCamp` 部分包装在 `a` 元素中，将 `href` 设置为 `https://freecodecamp.org`，并将 `target` 设置为 `_blank`。


在下面添加第二个 `p` 元素，类为 `publish-date`，文本为 `March 7, 2019`。

# --hints--

你应该创建一个新的 `div` 元素。

```js
assert.exists(document.querySelector('div'));
```

你的 `div` 元素应该位于 `header` 元素之后。

```js
assert.equal(document.querySelector('div')?.previousElementSibling?.localName, 'header');
```

你的 `div` 元素应将 `class` 设置为 `author`。

```js
assert.equal(document.querySelector('div')?.className, 'author');
```

你应该创建两个新的 `p` 元素。

```js
assert.lengthOf(document.querySelectorAll('p'), 3);
```

你的两个新的 `p` 元素应该位于 `div`元素中。

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

你的第一个新的 `p` 元素应该将 `class` 设置为 `author-name`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.className, 'author-name');
```

你的第一个新的 `p` 元素应该包含文本 `By freeCodeCamp`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.innerText, 'By freeCodeCamp');
```

你的第二个新的 `p` 元素应将 `class` 设置为 `publish-date`。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className, 'publish-date');
```

你第二个新的 `p` 元素应该包含文本 `March 7, 2019`。

```js
assert.equal(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText, 'March 7, 2019');
```

你应该创建一个新的 `a` 元素。

```js
assert.exists(document.querySelector('a'));
```

`a` 元素应该包含在你第一个新的 `p` 元素之内。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName, 'a');
```

你的 `a` 元素应该将 `href` 设置为 `https://freecodecamp.org`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href'), 'https://freecodecamp.org');
```

你的 `a` 元素应将 `target` 设置为 `_blank`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target'), '_blank');
```

你的 `a` 元素应包含文本 `freeCodeCamp`。

```js
assert.equal(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText, 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
