---
id: 614385513d91ae5c251c2052
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在这个工作坊，你将使用 freeCodeCamp 课程的前一版本的新闻稿创建一个杂志页面。

首先为 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap` 字体样式表添加一个 `link` 元素，为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css` FontAwesome 样式表添加一个 `link`，并为你的 `./styles.css` 样式表添加一个 `link`。

你的字体样式表将加载三种不同的字体：`Anton`、`Baskervville` 和 `Raleway`。

# --hints--

你的代码应该包含三个 `link` 元素。

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

你的 `link` 元素应该在你的 `head` 元素中。

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

你的三个 `link` 元素应该有一个 `rel` 属性，其值为 `stylesheet`。

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

你的一个链接元素的 `href` 应设置为 `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

你的链接元素之一应将 `href` 设置为 `https://use.fontawesome.com/releases/v5.8.2/css/all.css`。

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

有一个 `link` 元素应将 `href` 应设置为 `styles.css`。

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
