---
id: a5deed1811a43193f9f1c841
title: 요소 건너뛰기 함수 구현하기
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

이 실습에서는 특정 테스트 함수에 따라 허용 가능한 요소를 찾을 때까지 배열의 요소를 건너뛰는 함수를 만드세요.

예를 들어, `[1, 1, 1, 2, 1, 1, 1]` 같은 배열과 `function(n) {return n === 2}` 같은 테스트 함수가 있을 때, 이 함수에 허용되는 첫 번째 요소는 인덱스 3에 있으므로 그 이전의 모든 요소는 버려야 하며, 출력은 나머지 요소 `[2, 1, 1, 1]`가 되어야 합니다.

**목표**: 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리**

1. `dropElements` 함수가 배열(`arr`)과 함수(`func`)를 인수로 받아야 합니다.
1. `dropElements` 함수는 배열을 순차적으로 반복하며 `func`이 요소에 대해 `true`을 반환할 때까지 첫 번째 요소부터 제거해야 합니다.
1. `dropElements` 함수는 조건이 충족되면 배열의 나머지 요소를 반환해야 합니다.
1. 조건이 한 번도 충족되지 않으면 빈 배열을 반환해야 합니다.

# --hints--

`dropElements` 함수가 있어야 합니다.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})`은 `[3, 4]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})`은 `[1, 0, 1]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})`은 `[1, 2, 3]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})`은 `[]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})`은 `[7, 4]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})`은 `[3, 9, 2]`를 반환해야 합니다.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
