---
id: 66db529d37ad966480ebb633
title: Construir un Programa Selector de Almuerzo
challengeType: 26
dashedName: build-a-lunch-picker-program
---

# --description--

En este laboratorio, construirás un programa que ayuda a gestionar las opciones de almuerzo. Trabajarás con un array de almuerzos, agregando y eliminando elementos del array, y seleccionando aleatoriamente una opción de almuerzo.

**Objetivo:** Cumple con las historias de usuario a continuación y pasa todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. Deberías crear una variable `lunches` y asignarla a un array vacío que se utilizará para almacenar elementos del almuerzo.

2. Deberías crear una función `addLunchToEnd` que reciba un array como primer argumento y una cadena como segundo argumento. La función debería:
    - Agregar la cadena al final del array.
    - Registrar la cadena `"[Lunch Item] agregado al final del menú de almuerzo."` en la consola, donde `[Lunch Item]` es la cadena pasada a la función.
    - Devuelve el array actualizado.

3. Deberías crear una función `addLunchToStart` que reciba un array como primer argumento y una cadena como segundo argumento. La función debería:
    - Agregar la cadena al inicio del array.
    - Registrar la cadena `"[Lunch Item] agregado al inicio del menú de almuerzo."` en la consola, donde `[Lunch Item]` es la cadena pasada a la función.
    - Devuelve el array actualizado.

4. Deberías crear una función `removeLastLunch` que reciba un array como argumento. La función debería:
   - Eliminar el último elemento del array.
   - Si la eliminación tiene éxito, registra la cadena `"[Lunch Item] eliminado del final del menú de almuerzo."` en la consola, donde `[Lunch Item]` es el elemento eliminado del array.
   - Si el array está vacío, registra la cadena `"No hay almuerzos para eliminar."` en la consola.
   - Devuelve el array actualizado.

5. Deberías crear una función `removeFirstLunch` que reciba un array como argumento. La función debería:
   - Eliminar el primer elemento del array.
   - Si la eliminación tiene éxito, registra la cadena `"[Lunch Item] eliminado del inicio del menú de almuerzo."` en la consola, donde `[Lunch Item]` es el elemento eliminado del array.
   - Si el array está vacío, registra la cadena `"No hay almuerzos para eliminar."` en la consola.
   - Devuelve el array actualizado.

6. Deberías crear una función `getRandomLunch` que reciba un array como argumento. La función debería:
   - Seleccionar aleatoriamente un elemento del array.
   - Si tiene éxito, registra la cadena `"Almuerzo seleccionado aleatoriamente: [Lunch Item]"` en la consola, donde `[Lunch Item]` es un elemento aleatorio en el array.
   - Si el array está vacío, registra la cadena `"No hay almuerzos disponibles."` en la consola.

7. Deberías crear una función `showLunchMenu` que reciba un array como argumento y:
   - Si hay elementos en el array, registra la cadena `"Elementos del menú: [Lunch Item], [Lunch Item]..."` en la consola, donde cada `[Lunch item]` es uno de los elementos en el array, en orden.
   - Si el array está vacío, registra la cadena `"El menú está vacío."` en la consola.

# --hints--

Deberías declarar una variable `lunches` y asignarla a un array vacío para almacenar los elementos del almuerzo.

```js
assert.isArray(lunches);
```

Deberías definir una función `addLunchToEnd`.

```js
assert.isFunction(addLunchToEnd);
```

La función `addLunchToEnd` debería tener dos parámetros.

```js
assert.lengthOf(addLunchToEnd, 2);
```

`addLunchToEnd(lunches, "Tacos")` debería registrar la cadena `"Tacos agregado al final del menú de almuerzo."` en la consola.

```js
const tempArr = [];
const temp = console.log;
const testLunches = []
try {
  console.log = obj => tempArr.push(obj);
  addLunchToEnd(testLunches, "Tacos");
  assert.strictEqual(tempArr[0], "Tacos added to the end of the lunch menu.");
  addLunchToEnd(testLunches, "Pizza");
  assert.strictEqual(tempArr[1], "Pizza added to the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToEnd(["Pizza", "Tacos"], "Burger")` debería devolver `["Pizza", "Tacos", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToEnd(["Pizza", "Tacos"], "Burger"), ["Pizza", "Tacos", "Burger"]);
  assert.deepStrictEqual(addLunchToEnd(["Fries"], "Tacos"), ["Fries", "Tacos"]);
} finally {
  console.log = temp;
}
```

Deberías definir una función `addLunchToStart`.

```js
assert.isFunction(addLunchToStart);
```

La función `addLunchToStart` debería tener dos parámetros.

```js
assert.lengthOf(addLunchToStart, 2);
```

`addLunchToStart(lunches, "Sushi")` debería registrar la cadena `"Sushi agregado al inicio del menú de almuerzo."` en la consola.

```js
const tempArr = [];
const testLunches = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  addLunchToStart(testLunches, "Sushi");
  assert.strictEqual(tempArr[0], "Sushi added to the start of the lunch menu.");  
  addLunchToStart(testLunches, "Burger");
  assert.strictEqual(tempArr[1], "Burger added to the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`addLunchToStart(["Burger", "Sushi"], "Pizza")` debería devolver `["Pizza", "Burger", "Sushi"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(addLunchToStart(["Burger", "Sushi"], "Pizza"), ["Pizza", "Burger", "Sushi"]);
  assert.deepStrictEqual(addLunchToStart(["Noodles"], "Sushi"), ["Sushi", "Noodles"]);
} finally {
  console.log = temp;
}
```

Deberías definir una función `removeLastLunch`.

```js
assert.isFunction(removeLastLunch)
```

La función `removeLastLunch` debería tener un parámetro.

```js
assert.lengthOf(removeLastLunch, 1);
```

Cuando el array de entrada está vacío, la función `removeLastLunch` debería registrar la cadena `"No hay almuerzos para eliminar."` en la consola.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Stew", "Soup", "Toast"])` debería registrar la cadena `"Toast eliminado del final del menú de almuerzo."` en la consola.

```js
let testLunches = ["Stew", "Soup", "Toast"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[0], "Toast removed from the end of the lunch menu.");
  testLunches = ["Rice", "Pizza"];
  removeLastLunch(testLunches);
  assert.strictEqual(tempArr[1], "Pizza removed from the end of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeLastLunch(["Sushi", "Pizza", "Noodles"])` debería devolver `["Sushi", "Pizza"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeLastLunch(["Sushi", "Pizza", "Noodles"]), ["Sushi", "Pizza"]);
  assert.deepStrictEqual(removeLastLunch(["Pizza", "Burger"]), ["Pizza"]);
} finally {
  console.log = temp;
}
```

Deberías definir una función `removeFirstLunch`.

```js
assert.isFunction(removeFirstLunch);
```

La función `removeFirstLunch` debería tener un solo parámetro.

```js
assert.lengthOf(removeFirstLunch, 1);
```

Cuando el array de entrada está vacío, la función `removeFirstLunch` debería registrar la cadena `"No hay almuerzos para eliminar."` en la consola.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches to remove.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Salad", "Eggs", "Cheese"])` debería registrar la cadena `"Salad eliminado del inicio del menú de almuerzo."` en la consola.

```js
let testLunches = ["Salad", "Eggs", "Cheese"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[0], "Salad removed from the start of the lunch menu.");
  testLunches = ["Stew", "Soup", "Toast"]
  removeFirstLunch(testLunches);
  assert.strictEqual(tempArr[1], "Stew removed from the start of the lunch menu.");
} finally {
  console.log = temp;
}
```

`removeFirstLunch(["Sushi", "Pizza", "Burger"])` debería devolver `["Pizza", "Burger"]`.

```js
const temp = console.log;
console.log = () => {}
try {
  assert.deepStrictEqual(removeFirstLunch(["Sushi", "Pizza", "Burger"]), ["Pizza", "Burger"]);
  assert.deepStrictEqual(removeFirstLunch(["Pizza", "Burger"]), ["Burger"]);
} finally {
  console.log = temp;
}
```

`addLunchToEnd`, `addLunchToStart`, `removeLastLunch`, y `removeFirstLunch` deberían devolver el mismo array pasado como argumento después de actualizarlo.

```js
const temp = console.log;
console.log = () => {}
const testLunches = ["Spam"];
try {
  assert.equal(addLunchToEnd(testLunches, "Pizza"), testLunches);
  assert.equal(addLunchToStart(testLunches, "Caviar"), testLunches);
  assert.equal(removeLastLunch(testLunches), testLunches);
  assert.equal(removeFirstLunch(testLunches), testLunches);
} finally {
  console.log = temp;
}
```

Deberías definir una función `getRandomLunch`.

```js
assert.isFunction(getRandomLunch);
```

La función `getRandomLunch` debería tener un solo parámetro.

```js
assert.lengthOf(getRandomLunch, 1);
```

Cuando el array de entrada está vacío, la función `getRandomLunch` debería registrar la cadena `"No hay almuerzos disponibles."` en la consola.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  getRandomLunch(testLunches);
  assert.strictEqual(tempArr[0], "No lunches available.");
} finally {
  console.log = temp;
}
```

Cuando el array de entrada no está vacío, la función `getRandomLunch` debería registrar una cadena en el formato `"Almuerzo seleccionado aleatoriamente: [Lunch Item]"` en la consola.

```js
const testLunches = ["Fish", "Fries", "Roast"];
const tempRandom = Math.random;
const tempArr = [];
const temp = console.log;

try {
  console.log = obj => tempArr.push(obj);
  // check that it correctly outputs the first item
  Math.random = () => 0;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[0], `Randomly selected lunch: ${testLunches[0]}`);

  // Second item
  Math.random = () => 0.4;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[1], `Randomly selected lunch: ${testLunches[1]}`);

  // Third item
  Math.random = () => 0.8;
  getRandomLunch([...testLunches]);
  assert.strictEqual(tempArr[2], `Randomly selected lunch: ${testLunches[2]}`);

} finally {
  Math.random = tempRandom;
  console.log = temp;
}
```

La función `getRandomLunch` no debería cambiar el array pasado como argumento..

```js
const testLunches = ["Fish", "Fries", "Roast"];
let originalLunches = [...testLunches];
const tempRandom = Math.random;

try {
  Math.random = () => 0;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.4;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  Math.random = () => 0.8;
  getRandomLunch(testLunches);
  assert.sameOrderedMembers(testLunches, originalLunches);

  const emptyLunches = [];
  let originalEmpty = [...emptyLunches];
  getRandomLunch(emptyLunches);
  assert.sameOrderedMembers(emptyLunches, originalEmpty);

} finally {
  Math.random = tempRandom;
}
```

Deberías definir una función `showLunchMenu`.

```js
assert.isFunction(showLunchMenu);
```

La función `showLunchMenu` debería tener un solo parámetro.

```js
assert.lengthOf(showLunchMenu, 1);
```

Cuando el array de entrada está vacío, la función `showLunchMenu` debería registrar la cadena `"El menú está vacío."` en la consola.

```js
const testLunches = [];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], "The menu is empty.");
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Greens", "Corns", "Beans"])` debería registrar `"Elementos del menú: Greens, Corns, Beans"` en la consola.

```js
const testLunches = ["Greens", "Corns", "Beans"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

`showLunchMenu(["Pizza", "Burger", "Fries", "Salad"])` debería registrar `"Elementos del menú: Pizza, Burger, Fries, Salad"` en la consola.

```js
const testLunches = ["Pizza", "Burger", "Fries", "Salad"];
const tempArr = [];
const temp = console.log;
try {
  console.log = obj => tempArr.push(obj);
  showLunchMenu(testLunches);
  assert.strictEqual(tempArr[0], `Menu items: ${testLunches.join(", ")}`);
} finally {
  console.log = temp;
}
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const lunches = [];

// Add a new lunch to the end of the list
function addLunchToEnd(lunches, newLunch) {
  lunches.push(newLunch);
  console.log(`${newLunch} added to the end of the lunch menu.`);
  return lunches;
}

// Add a new lunch to the start of the list
function addLunchToStart(lunches, newLunch) {
  lunches.unshift(newLunch);
  console.log(`${newLunch} added to the start of the lunch menu.`);
  return lunches;
}

// Remove the last lunch from the list
function removeLastLunch(lunches) {
  const removedLunch = lunches.pop();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the end of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Remove the first lunch from the list
function removeFirstLunch(lunches) {
  const removedLunch = lunches.shift();
  if (removedLunch) {
    console.log(`${removedLunch} removed from the start of the lunch menu.`);
  } else {
    console.log("No lunches to remove.");
  }
  return lunches;
}

// Function to get a random lunch from the list
function getRandomLunch(lunches) {
  if (lunches.length === 0) {
    console.log("No lunches available.");
  }
  else {
    const randomIndex = Math.floor(Math.random() * lunches.length);
    const randomLunch = lunches[randomIndex];
    console.log(`Randomly selected lunch: ${randomLunch}`);
  }
}

// Function to display all the lunches in the list
function showLunchMenu(lunches) {
  if (lunches.length === 0) {
    console.log("The menu is empty.");
  }
  else {
    console.log(`Menu items: ${lunches.join(', ')}`);
  }
}

addLunchToEnd(lunches, "Tacos");
addLunchToStart(lunches, "Sushi");
removeLastLunch(lunches);
removeFirstLunch(lunches);
getRandomLunch(lunches);
showLunchMenu(lunches);

```
