---
id: 641da5462576784453146ec2
title: Крок 7
challengeType: 0
dashedName: step-7
---

# --description--

Тепер, коли у вас є потрібні вам дані, можна використати їх для заповнення інтерфейсу користувача. Але отримані дані містять масив із 26 авторів, і якщо додати їх усіх на сторінку одночасно, це може призвести до низької продуктивності.

Натомість ви повинні додати 8 авторів за раз та мати кнопку для додавання ще 8, доки не залишиться даних для показу.

Використайте `let`, щоб створити 2 змінні з іменами `startingIndex` та `endingIndex` і присвойте їм числові значення `0` та `8` відповідно. Крім того, створіть змінну `authorDataArr` з `let` та встановіть її на порожній масив.

# --before-all--

```js
window.fetch = () => Promise.resolve({json: () => Promise.resolve([{ author: 'Whoever', image: 'data:image/gif;base64,R0lGODlhAQABAIAAAAAAAP///ywAAAAAAQABAAACAUwAOw==', url: "http://not-a-real-url.nowhere/", bio: 'words go here' }])});
```

# --hints--

Використайте `let`, щоб оголосити змінну під назвою `startingIndex`.

```js
assert.match(code, /let\s+startingIndex/)
```

Встановіть змінну `startingIndex` на `0`.

```js
assert.match(code, /let\s+startingIndex\s*=\s*0\s*;?/)
```

Використайте `let`, щоб оголосити змінну під назвою `endingIndex`.

```js
assert.match(code, /let\s+endingIndex/)
```

Встановіть змінну `endingIndex` на `8`.

```js
assert.match(code, /let\s+endingIndex\s*=\s*8\s*;?/)
```

Використайте `let`, щоб оголосити змінну під назвою `authorDataArr`.

```js
assert.match(code, /let\s+authorDataArr/)
```

Встановіть змінну `authorDataArr` на порожній масив (`[]`).

```js
assert.match(code, /let\s+authorDataArr\s*=\s*\[\s*\]\s*;?/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta http-equiv="X-UA-Compatible" content="IE=edge" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>freeCodeCamp News Author Page</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <h1 class="title">freeCodeCamp News Author Page</h1>

    <main>
      <div id="author-container"></div>
      <button class="btn" id="load-more-btn">Load More Authors</button>
    </main>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
* {
  margin: 0;
  padding: 0;
  box-sizing: border-box;
}

:root {
  --main-bg-color: #1b1b32;
  --light-grey: #f5f6f7;
  --dark-purple: #5a01a7;
  --golden-yellow: #feac32;
}

body {
  background-color: var(--main-bg-color);
  text-align: center;
}

.title {
  color: var(--light-grey);
  margin: 20px 0;
}

#author-container {
  display: flex;
  flex-wrap: wrap;
  justify-content: center;
}

.user-card {
  border-radius: 15px;
  width: 300px;
  height: 350px;
  background-color: var(--light-grey);
  margin: 20px;
}

.user-img {
  width: 150px;
  height: 150px;
  object-fit: cover;
}

.purple-divider {
  background-color: var(--dark-purple);
  width: 100%;
  height: 15px;
}

.author-name {
  margin: 10px;
}

.bio {
  margin: 20px;
}

.error-msg {
  color: var(--light-grey);
}

.btn {
  cursor: pointer;
  width: 200px;
  margin: 10px;
  color: var(--main-bg-color);
  font-size: 14px;
  background-color: var(--golden-yellow);
  background-image: linear-gradient(#fecc4c, #ffac33);
  border-color: var(--golden-yellow);
  border-width: 3px;
}
```

```js
const authorContainer = document.getElementById('author-container');
const loadMoreBtn = document.getElementById('load-more-btn');

--fcc-editable-region--

--fcc-editable-region--

fetch('https://cdn.freecodecamp.org/curriculum/news-author-page/authors.json')
  .then((res) => res.json())
  .then((data) => {
    console.log(data);   
  })
  .catch((err) => {
    console.error(`There was an error: ${err}`);
  });
```
