---
id: 6a57c331ae94fa4631b60866
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Às vezes a equipe vai pedir um artefato que o museu não tem. Se nenhum artefato estiver armazenado sob aquele `id`, `collection[id]` será `undefined` e ler `.title` dele causaria um erro.

Atualize a instrução return para usar uma condicional. Retorne `artifact.title` se um artefato for encontrado e a string `"Artifact not found"` se não for.

# --hints--

`getArtifactTitle(101)` ainda deve retornar `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle` deve retornar `"Artifact not found"` se nenhum artefato corresponder ao `id`.

```js
assert.strictEqual(getArtifactTitle(999), "Artifact not found");
```

`getArtifactTitle` deve retornar o título de qualquer artefato que encontrar, não apenas o `101`.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
assert.strictEqual(getArtifactTitle(457), "Test Artifact");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
--fcc-editable-region--
  return artifact.title;
--fcc-editable-region--
}
```
