---
id: 6a57c331ae94fa4631b6086f
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Atualize o curador do artefato `101` para `"Fran Sinclair"`. Depois, faça uma chamada a `console.log()` com `collection[101].curator.name` para registrar o nome do curador desse artefato.

# --hints--

Você deve fazer uma chamada a `updateCurator` com `101` e `"Fran Sinclair"`.

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)/);
```

Você deve registrar o `name` do curador do primeiro artefato novamente, depois de atualizá-lo.

```js
assert.match(__helpers.removeJSComments(code), /updateCurator\s*\(\s*101\s*,\s*["']Fran Sinclair["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*curator|\[\s*["']curator["']\s*\])(?:\.\s*name|\[\s*["']name["']\s*\])\s*\)/);
```

O `name` do curador do primeiro artefato deve ser `"Fran Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Fran Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

function moveArtifact(id, gallery, year) {
  const artifact = collection[id];
  if (artifact) {
    artifact.locations.push({ gallery, year });
  }
}

moveArtifact(102, "Hall B", 2026);
console.log(collection[102].locations);

function toggleDisplayStatus(id) {
  const artifact = collection[id];
  if (artifact) {
    artifact.onDisplay = !artifact.onDisplay;
  }
}

console.log(collection[102].onDisplay);
toggleDisplayStatus(102);
console.log(collection[102].onDisplay);

function updateCurator(id, name) {
  const artifact = collection[id];
  if (artifact) {
    artifact.curator.name = name;
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
