---
id: 6a57c331ae94fa4631b60868
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Curadores adicionam tags aos artefatos para facilitar a busca, e uma tag nunca deve ser adicionada duas vezes. Antes de adicionar uma nova tag, você pode usar o método `includes()` para verificar se ela já existe, pois ele retorna `true` quando um array já contém um valor.

Crie uma função chamada `addTag` que receba os parâmetros `id` e `tag`. Procure o artefato usando `collection[id]`. Se ele existir e seu array `tags` ainda não incluir `tag`, adicione `tag` ao `artifact.tags`.

# --hints--

Você deve ter uma função chamada `addTag` com os parâmetros chamados `id` e `tag`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.addTag);
const params = explorer.allFunctions.addTag.parameters;
assert.equal(params.length, 2);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "tag");
```

Chamar `addTag` com uma nova tag deve adicionar essa tag ao array `tags` do artefato.

```js
addTag(102, "ancient");
assert.include(collection[102].tags, "ancient");
```

Chamar `addTag` com uma tag que o artefato já possui não deve adicionar um duplicado.

```js
const tagsBefore = [...collection[102].tags];
addTag(102, "bronze");
assert.deepEqual(collection[102].tags, tagsBefore);

const start = collection[102].tags.length;
addTag(102, "unique-check");
addTag(102, "unique-check");
assert.strictEqual(collection[102].tags.length, start + 1);
```

`addTag` deve adicionar a tag ao artefato ao qual o `id` se refere.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
addTag(457, "test-tag");
assert.deepEqual(collection[457].tags, ["test-tag"]);
```

`addTag` não deve lançar um erro nem alterar a coleção quando nenhum artefato corresponder ao `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => addTag(999, "test"));
assert.strictEqual(JSON.stringify(collection), before);
```

Você deve adicionar `tag` ao array `tags` existente do artefato, em vez de substituir o array.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const tagsBefore = collection[458].tags;
const spy = __helpers.spyOn(tagsBefore, "push");
try {
  addTag(458, "pushed");
  assert.deepEqual(spy.calls, [["pushed"]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].tags, tagsBefore);
assert.deepEqual(collection[458].tags, ["pushed"]);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

--fcc-editable-region--

--fcc-editable-region--
```
