---
id: 6a57c331ae94fa4631b60865
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

A equipe consulta artefatos dezenas de vezes por dia, então seu auditor precisa de uma busca reutilizável. Os colchetes também podem receber uma variável, então `collection[id]` retorna o artefato ao qual aquele `id` se refere.

Crie uma função chamada `getArtifactTitle` que receba um parâmetro `id`. Use `collection[id]` para obter o artefato, armazene-o em uma variável chamada `artifact` e retorne `artifact.title`.

# --hints--

Você deve ter uma função chamada `getArtifactTitle` com um parâmetro chamado `id`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle);
const params = explorer.allFunctions.getArtifactTitle.parameters;
assert.equal(params.length, 1);
assert.equal(params[0].toString(), "id");
```

Você deve atribuir `collection[id]` a uma variável chamada `artifact`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.getArtifactTitle?.variables.artifact);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.variables.artifact?.value.matches("collection[id]"));
```

`getArtifactTitle` deve retornar `artifact.title`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.allFunctions.getArtifactTitle?.hasReturn("artifact.title"));
```

`getArtifactTitle(101)` deve retornar `"Golden Mask"`.

```js
assert.strictEqual(getArtifactTitle(101), "Golden Mask");
```

`getArtifactTitle(102)` deve retornar `"Bronze Tablet"`.

```js
assert.strictEqual(getArtifactTitle(102), "Bronze Tablet");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

--fcc-editable-region--

--fcc-editable-region--
```
