---
id: 6a57c331ae94fa4631b60869
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Novas pesquisas confirmam que a máscara dourada pertencia à realeza. Adicione a tag `"royal"` ao artefato com `id` `101`. Depois, faça uma chamada a `console.log()` com `collection[101].tags` para registrar as tags desse artefato.

# --hints--

Você deve chamar `addTag` com `101` e `"royal"`.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)/);
```

Você deve registrar as `tags` do primeiro artefato no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

Você deve registrar as `tags` depois de adicionar a nova tag.

```js
assert.match(__helpers.removeJSComments(code), /addTag\s*\(\s*101\s*,\s*["']royal["']\s*\)[\s\S]*console\.log\s*\(\s*collection\s*\[\s*["']?101["']?\s*\](?:\.\s*tags|\[\s*["']tags["']\s*\])\s*\)/);
```

As `tags` do primeiro artefato devem incluir `"royal"`.

```js
assert.include(collection[101].tags, "royal");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

--fcc-editable-region--

--fcc-editable-region--
```
