---
id: 6a57c265abbc95dfebc05098
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai construir um auditor de proveniência de artefatos. Você vai armazenar artefatos de museu em um objeto e escrever funções para consultá-los, etiquetá-los, movê-los entre galerias e resumir sua história.

Cada artefato é armazenado sob seu id. Crie um objeto chamado `collection` e adicione uma entrada a ele. Defina a chave como `101` e o valor como um objeto com um `title` de `"Golden Mask"` e uma `category` de `"Ceremonial"`.

# --hints--

Você deve declarar uma variável chamada `collection`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.variables.collection);
```

`collection` deve ser um objeto.

```js
assert.isObject(collection);
```

`collection` deve ter uma entrada.

```js
assert.strictEqual(Object.keys(collection).length, 1);
```

O artefato armazenado sob `101` deve ser um objeto.

```js
assert.isObject(collection[101]);
```

O artefato armazenado sob `101` deve ter um `title` de `"Golden Mask"`.

```js
assert.strictEqual(collection[101].title, "Golden Mask");
```

O artefato armazenado sob `101` deve ter uma `category` de `"Ceremonial"`.

```js
assert.strictEqual(collection[101].category, "Ceremonial");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
