---
id: 6a57c331ae94fa4631b6086a
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

À medida que os artefatos se movem entre galerias ao longo do tempo, cada movimento é registrado como uma nova entrada no array `locations`.

Crie uma função `moveArtifact` com os parâmetros `id`, `gallery` e `year`. Busque o artefato usando `collection[id]`. Se ele existir, adicione um novo objeto com `gallery` e `year` ao seu array `locations`.

# --hints--

Você deve ter uma função chamada `moveArtifact` com os parâmetros `id`, `gallery` e `year`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.moveArtifact);
const params = explorer.allFunctions.moveArtifact.parameters;
assert.equal(params.length, 3);
assert.equal(params[0].toString(), "id");
assert.equal(params[1].toString(), "gallery");
assert.equal(params[2].toString(), "year");
```

Chamar `moveArtifact` deve adicionar um novo objeto `{ gallery, year }` ao array `locations` do artefato.

```js
const start = collection[101].locations.length;
moveArtifact(101, "Test Hall", 2030);
assert.lengthOf(collection[101].locations, start + 1);
assert.deepEqual(collection[101].locations[collection[101].locations.length - 1], { gallery: "Test Hall", year: 2030 });
```

`moveArtifact` deve adicionar a nova localização ao artefato ao qual o `id` se refere.

```js
collection[457] = {
  title: "Test Artifact",
  category: "Test Category",
  curator: { id: 457, name: "Test Curator" },
  locations: [{ gallery: "Hall Y", year: 2001 }],
  tags: [],
  onDisplay: false,
};
moveArtifact(457, "Test Vault", 2032);
assert.lengthOf(collection[457].locations, 2);
assert.deepEqual(collection[457].locations[1], { gallery: "Test Vault", year: 2032 });
```

`moveArtifact` não deve lançar um erro nem alterar a coleção quando nenhum artefato corresponder ao `id`.

```js
const before = JSON.stringify(collection);
assert.doesNotThrow(() => moveArtifact(999, "Test Hall", 2030));
assert.strictEqual(JSON.stringify(collection), before);
```

Você deve adicionar a nova localização ao array `locations` existente do artefato, em vez de substituir o array.

```js
collection[458] = {
  title: "Push Test",
  category: "Test Category",
  curator: { id: 458, name: "Test Curator" },
  locations: [{ gallery: "Hall Z", year: 2002 }],
  tags: [],
  onDisplay: false,
};
const locationsBefore = collection[458].locations;
const spy = __helpers.spyOn(locationsBefore, "push");
try {
  moveArtifact(458, "Push Hall", 2033);
  assert.deepEqual(spy.calls, [[{ gallery: "Push Hall", year: 2033 }]]);
} finally {
  spy.restore();
}
assert.strictEqual(collection[458].locations, locationsBefore);
assert.deepEqual(collection[458].locations[1], { gallery: "Push Hall", year: 2033 });
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
    locations: [
      { gallery: "Hall A", year: 2020 },
      { gallery: "Hall C", year: 2024 },
    ],
    tags: ["gold", "egypt"],
    onDisplay: true,
  },
  102: {
    title: "Bronze Tablet",
    category: "Inscription",
    curator: {
      id: 202,
      name: "Robert Sinclair",
    },
    locations: [{ gallery: "Archive Wing", year: 2019 }],
    tags: ["bronze", "writing"],
    onDisplay: false,
  },
};

console.log(collection[101].title);
console.log(collection[101].curator.name);

function getArtifactTitle(id) {
  const artifact = collection[id];
  return artifact ? artifact.title : "Artifact not found";
}

console.log(getArtifactTitle(102));

function addTag(id, tag) {
  const artifact = collection[id];
  if (artifact && !artifact.tags.includes(tag)) {
    artifact.tags.push(tag);
  }
}

addTag(101, "royal");
console.log(collection[101].tags);

--fcc-editable-region--

--fcc-editable-region--
```
