---
id: 6a57c331ae94fa4631b60861
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Todo artefato tem uma pessoa responsável por ele.

Adicione uma propriedade `curator` ao artefato armazenado sob a chave `101` e defina seu valor como um objeto com um `id` de `201` e um `name` de `"Earl Sinclair"`.

# --hints--

O artefato armazenado sob `101` deve ter uma propriedade `curator`.

```js
assert.property(collection[101], "curator");
```

O `curator` deve ser um objeto.

```js
assert.isObject(collection[101].curator);
```

O `curator` deve ter um `id` de `201`.

```js
assert.strictEqual(collection[101].curator.id, 201);
```

O `curator` deve ter um `name` igual a `"Earl Sinclair"`.

```js
assert.strictEqual(collection[101].curator.name, "Earl Sinclair");
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
