---
id: 6a57c331ae94fa4631b60862
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Cada artefato também registra onde esteve, como é descrito e se está em exibição.

Adicione mais três propriedades ao artefato armazenado sob `101`:

- `locations`: um array contendo dois objetos, `{ gallery: "Hall A", year: 2020 }` e `{ gallery: "Hall C", year: 2024 }`
- `tags`: um array contendo as strings `"gold"` e `"egypt"`
- `onDisplay`: o booleano `true`

# --hints--

O `locations` do artefato deve ser um array dos dois objetos de localização fornecidos.

```js
assert.deepEqual(collection[101].locations, [
  { gallery: "Hall A", year: 2020 },
  { gallery: "Hall C", year: 2024 },
]);
```

O `tags` do artefato deve ser `["gold", "egypt"]`.

```js
assert.deepEqual(collection[101].tags, ["gold", "egypt"]);
```

O `onDisplay` do artefato deve ser `true`.

```js
assert.strictEqual(collection[101].onDisplay, true);
```

# --seed--

## --seed-contents--

```js
const collection = {
  101: {
    title: "Golden Mask",
    category: "Ceremonial",
    curator: {
      id: 201,
      name: "Earl Sinclair",
    },
--fcc-editable-region--
    
--fcc-editable-region--
  },
};
```
