---
id: 587d7dbd367417b2b2512bb6
title: 用 Mixins 创建可重用 CSS
challengeType: 0
forumTopicId: 301455
dashedName: create-reusable-css-with-mixins
---

# --description--

In Sass, a <dfn>mixin</dfn> is a group of CSS declarations that can be reused throughout the style sheet. The definition starts with the `@mixin` at-rule, followed by a custom name. You apply the mixin using the `@include` at-rule.

```scss
@mixin reset-list {
  margin: 0;
  padding: 0;
  list-style: none;
}

nav ul {
  @include reset-list;
}
```

编译为：

```css
nav ul {
  margin: 0;
  padding: 0;
  list-style: none;
}
```

你的 mixin 还可以接受参数，从而允许定制它们的行为。 使用 mixin 时需要参数。

```scss
@mixin prose($font-size, $spacing) {
  font-size: $font-size;
  margin: 0;
  margin-block-end: $spacing;
}

p {
  @include prose(1.25rem, 1rem);
}

h2 {
  @include prose(2.4rem, 1.5rem);
}
```

你可以通过赋予参数默认值来使参数成为可选的。

```scss
@mixin text-color($color: black) {
  color: $color;
}

p {
  @include text-color(); /* color: black */
}

nav a {
  @include text-color(orange);
}
```

# --instructions--

编写一个名为 `shape` 的 mixin，并赋予其 3 个参数：`$w`、`$h` 和 `$bg-color`。

使用 `shape` 混合为 `#square` 元素提供 `50px` 的宽度和高度，以及 `red` 的背景颜色。 对于 `#rect-a` 元素，添加宽度 `100px`、高度 `50px` 和背景颜色 `blue`。 最后，为 `#rect-b` 元素添加宽度 `50px`、高度 `100px` 和背景颜色 `orange`。

# --hints--

你应该声明一个名为 `shape` 的 mixin，它有 3 个参数：`$w`、`$h` 和 `$bg-color`。

```js
assert.match(code, /@mixin\s+shape\s*\(\s*\$w,\s*\$h,\s*\$bg-color\s*\)\s*{/gi);
```

你的 mixin 应该包含一个使用 `$w` 参数的 `width` 属性。

```js
assert.match(__helpers.removeWhiteSpace(code), /width:\$w;/gi);
```

你的 mixin 应该包含一个使用 `$h` 参数的 `height` 属性。

```js
assert.match(__helpers.removeWhiteSpace(code), /height:\$h;/gi);
```

你的 mixin 应该包含一个使用 `$bg-color` 参数的 `background-color` 属性。

```js
assert.match(__helpers.removeWhiteSpace(code), /background-color:\$bg\-color;/gi);
```

你应该使用 `@include` 关键字将 `#square` 选择器内的样式替换为对 `shape` 混合的调用。 设置宽度和高度为 `50px`，背景颜色为 `red`。

```js
assert.match(code, /#square\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*50px\s*,\s*red\s*\)\s*;\s*}/gi);
```

你应该使用 `@include` 关键字将 `#rect-a` 选择器内的样式替换为对 `shape` 混合的调用。 设置宽度为 `100px`、高度为 `50px`、背景颜色为 `blue`。

```js
assert.match(code, /#rect-a\s*{\s*@include\s+shape\s*\(\s*100px\s*,\s*50px\s*,\s*blue\s*\)\s*;\s*}/gi);
```

你应该使用 `@include` 关键字将 `#rect-b` 选择器内的样式替换为对 `shape` 混合的调用。 设置宽度为 `50px`、高度为 `100px`、背景颜色为 `orange`。

```js
assert.match(code, /#rect-b\s*{\s*@include\s+shape\s*\(\s*50px\s*,\s*100px\s*,\s*orange\s*\)\s*;\s*}/gi);
```

# --seed--

## --seed-contents--

```html
<style type='text/scss'>
#square {
  width: 50px;
  height: 50px;
  background-color: red;
}

#rect-a {
  width: 100px;
  height: 50px;
  background-color: blue;
}

#rect-b {
  width: 50px;
  height: 100px;
  background-color: orange;
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```

# --solutions--

```html
<style type='text/scss'>
@mixin shape($w, $h, $bg-color) {
  width: $w;
  height: $h;
  background-color: $bg-color;
}

#square {
  @include shape(50px, 50px, red);
}

#rect-a {
  @include shape(100px, 50px, blue);
}

#rect-b {
  @include shape(50px, 100px, orange);
}
</style>

<div id="square"></div>
<div id="rect-a"></div>
<div id="rect-b"></div>
```
