---
id: 587d78ab367417b2b2512af0
title: 'Verwenden Sie display: flex, um zwei Kästen zu positionieren'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

Dieser Abschnitt verwendet abwechselnde Aufgabenstile, um zu zeigen, wie Sie CSS verwenden, um Elemente flexibel zu positionieren. Zuerst erklärt eine Aufgabe die Theorie, dann wendet eine praktische Aufgabe mit einer einfachen Tweet-Komponente das Flexbox-Konzept an.

Das Anwenden der CSS-Eigenschaft `display: flex;` auf ein Element ermöglicht es Ihnen, andere Flex-Eigenschaften zu verwenden, um eine responsiv gestaltete Seite zu erstellen.

# --instructions--

Fügen Sie die CSS-Eigenschaft `display` zu `#box-container` hinzu und setzen Sie ihren Wert auf `flex`.

# --hints--

`#box-container` sollte die Eigenschaft `display` mit dem Wert `flex` haben.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
