---
id: 587d78ad367417b2b2512afa
title: "Die Eigenschaft flex-wrap verwenden, um eine Zeile oder Spalte umzubrechen"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS Flexbox hat eine Funktion, mit der ein Flex-Container in mehrere Zeilen (oder Spalten) aufgeteilt werden kann. Standardmäßig passt ein Flex-Container alle Flex-Elemente zusammen. Zum Beispiel befindet sich eine Zeile mit allen Elementen auf einer Linie.

Wenn Sie jedoch die `flex-wrap`-Eigenschaft verwenden, teilt CSS die Elemente um. Das bedeutet, dass zusätzliche Elemente in eine neue Zeile oder Spalte verschoben werden. Der Haltepunkt, an dem der Umbruch erfolgt, hängt von der Größe der Elemente und der Größe des Containers ab.

CSS bietet auch Optionen für die Richtung des Umbruchs:

<ul><li><code>nowrap</code>: dies ist die Standardeinstellung und bricht die Elemente nicht um.</li><li><code>wrap</code>: bricht die Elemente auf mehrere Zeilen von oben nach unten um, wenn sie in Zeilen sind, und von links nach rechts, wenn sie in Spalten sind.</li><li><code>wrap-reverse</code>: bricht die Elemente auf mehrere Zeilen von unten nach oben um, wenn sie in Zeilen sind, und von rechts nach links, wenn sie in Spalten sind.</li></ul>

# --instructions--

Das aktuelle Layout hat zu viele Kästchen für eine Zeile. Fügen Sie dem `flex-wrap`-Element die CSS-Eigenschaft `#box-container` hinzu und geben Sie ihr den Wert `wrap`.

# --hints--

Das `#box-container`-Element sollte die Eigenschaft `flex-wrap` mit dem Wert `wrap` haben.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
