---
id: 587d78ae367417b2b2512afe
title: Die Kurzschreibweise-Eigenschaft flex verwenden
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

Es gibt eine Kurzschreibweise, mit der mehrere Flex-Eigenschaften gleichzeitig gesetzt werden können. Die Eigenschaften `flex-grow`, `flex-shrink` und `flex-basis` können alle zusammen mit der Eigenschaft `flex` gesetzt werden.

Zum Beispiel setzt `flex: 1 0 10px;` das Element auf `flex-grow: 1;`, `flex-shrink: 0;` und `flex-basis: 10px;`.

Die Standardeinstellungen der Eigenschaft sind `flex: 0 1 auto;`.

# --instructions--

Fügen Sie die CSS-Eigenschaft `flex` sowohl zu `#box-1` als auch zu `#box-2` hinzu. Geben Sie `#box-1` die Werte, sodass seine `flex-grow` `2` ist, seine `flex-shrink` `2` und seine `flex-basis` `150px` ist. Geben Sie `#box-2` die Werte, sodass seine `flex-grow` `1` ist, seine `flex-shrink` `1` und seine `flex-basis` `150px` ist.

Diese Werte bewirken, dass `#box-1` sich ausdehnt, um den zusätzlichen Platz mit der doppelten Geschwindigkeit von `#box-2` zu füllen, wenn der Container größer als 300px ist, und sich mit der doppelten Geschwindigkeit von `#box-2` verkleinert, wenn der Container kleiner als 300px ist. 300px ist die kombinierte Größe der `flex-basis` Werte der beiden Boxen.

# --hints--

Das `#box-1` Element sollte die Eigenschaft `flex` mit dem Wert `2 2 150px` haben.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

Das `#box-2` Element sollte die Eigenschaft `flex` mit dem Wert `1 1 150px` haben.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Ihr Code sollte die Eigenschaft `flex` für `#box-1` und `#box-2` verwenden.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
