---
id: 587d78af367417b2b2512b00
title: Die Eigenschaft align-self verwenden
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

Die letzte Eigenschaft für Flex-Elemente ist `align-self`. Diese Eigenschaft ermöglicht es Ihnen, die Ausrichtung jedes einzelnen Elements individuell anzupassen, anstatt sie alle auf einmal festzulegen. Das ist nützlich, da andere gängige Anpassungstechniken mit den CSS-Eigenschaften `float`, `clear` und `vertical-align` bei Flex-Elementen nicht funktionieren.

`align-self` akzeptiert dieselben Werte wie `align-items` und überschreibt jeden Wert, der durch die Eigenschaft `align-items` gesetzt wurde.

# --instructions--

Fügen Sie die CSS-Eigenschaft `align-self` sowohl zu `#box-1` als auch zu `#box-2` hinzu. Geben Sie `#box-1` den Wert `center` und `#box-2` den Wert `flex-end`.

# --hints--

Das `#box-1`-Element sollte die Eigenschaft `align-self` mit dem Wert `center` haben.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

Das `#box-2`-Element sollte die Eigenschaft `align-self` mit dem Wert `flex-end` haben.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
