---
id: 587d78ab367417b2b2512af2
title: "Die Eigenschaft flex-direction verwenden, um eine Zeile zu erstellen"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

Das Hinzufügen von `display: flex` zu einem Element macht es zu einem Flex-Container. Dadurch können alle Kinder dieses Elements in Reihen oder Spalten ausgerichtet werden. Dies erreichen Sie, indem Sie dem übergeordneten Element die `flex-direction`-Eigenschaft hinzufügen und sie auf row oder column setzen. Eine Zeile erzeugt eine horizontale Ausrichtung der Kinder, eine Spalte eine vertikale Ausrichtung.

Weitere Optionen für `flex-direction` sind `row-reverse` und `column-reverse`.

**Hinweis:** Der Standardwert für die `flex-direction`-Eigenschaft ist `row`.

# --instructions--

Fügen Sie dem `flex-direction`-Element die CSS-Eigenschaft `#box-container` hinzu und geben Sie ihr den Wert `row-reverse`.

# --hints--

Das `#box-container`-Element sollte die `flex-direction`-Eigenschaft mit dem Wert `row-reverse` haben.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
