---
id: 587d78ad367417b2b2512afb
title: "Die Eigenschaft flex-shrink verwenden, um Elemente zu verkleinern"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

Bisher beziehen sich alle Eigenschaften in den Aufgaben auf den Flex-Container (den übergeordneten Container der Flex-Elemente). Es gibt jedoch mehrere nützliche Eigenschaften für die Flex-Elemente.

Die erste ist die `flex-shrink`-Eigenschaft. Wenn sie verwendet wird, erlaubt sie einem Element, sich zu verkleinern, falls der Flex-Container zu klein ist. Elemente verkleinern sich, wenn die Breite des übergeordneten Containers kleiner ist als die kombinierte Breite aller Flex-Elemente darin.

Die `flex-shrink`-Eigenschaft nimmt Zahlen als Werte an. Je höher die Zahl, desto mehr wird das Element im Vergleich zu den anderen Elementen im Container schrumpfen. Zum Beispiel, wenn ein Element einen `flex-shrink`-Wert von `1` hat und das andere einen `flex-shrink`-Wert von `3`, wird das Element mit dem Wert `3` dreimal so stark schrumpfen wie das andere.

# --instructions--

Fügen Sie die CSS-Eigenschaft `flex-shrink` sowohl zu `#box-1` als auch zu `#box-2` hinzu. Geben Sie `#box-1` den Wert `1` und `#box-2` den Wert `2`.

# --hints--

Das `#box-1`-Element sollte die Eigenschaft `flex-shrink` mit dem Wert `1` haben.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

Das `#box-2`-Element sollte die Eigenschaft `flex-shrink` mit dem Wert `2` haben.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
