---
id: 587d78ae367417b2b2512aff
title: "Die Eigenschaft order verwenden, um Elemente neu anzuordnen"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

Die `order`-Eigenschaft wird verwendet, um CSS die Reihenfolge mitzuteilen, in der Flex-Elemente im Flex-Container angezeigt werden. Standardmäßig erscheinen die Elemente in der gleichen Reihenfolge, in der sie im Quell-HTML stehen. Die Eigenschaft nimmt Zahlen als Werte an, und es können auch negative Zahlen verwendet werden.

# --instructions--

Fügen Sie die CSS-Eigenschaft `order` sowohl zu `#box-1` als auch zu `#box-2` hinzu. Geben Sie `#box-1` den Wert `2` und `#box-2` den Wert `1`.

# --hints--

Das `#box-1`-Element sollte die `order`-Eigenschaft mit dem Wert `2` haben.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

Das `#box-2`-Element sollte die `order`-Eigenschaft mit dem Wert `1` haben.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
