---
id: 587d78ae367417b2b2512afc
title: "Die Eigenschaft flex-grow verwenden, um Elemente zu vergrößern"
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

Das Gegenteil von `flex-shrink` ist die Eigenschaft `flex-grow`. Erinnern Sie sich, dass `flex-shrink` die Größe der Elemente steuert, wenn der Container schrumpft. Die Eigenschaft `flex-grow` steuert die Größe der Elemente, wenn der übergeordnete Container wächst.

Wenn Sie ein ähnliches Beispiel aus der letzten Aufgabe verwenden: Hat ein Element einen `flex-grow`-Wert von `1` und das andere einen `flex-grow`-Wert von `3`, wächst das mit dem Wert `3` dreimal so stark wie das andere.

# --instructions--

Fügen Sie die CSS-Eigenschaft `flex-grow` sowohl zu `#box-1` als auch zu `#box-2` hinzu. Geben Sie `#box-1` den Wert `1` und `#box-2` den Wert `2`.

# --hints--

Das `#box-1`-Element sollte die Eigenschaft `flex-grow` mit dem Wert `1` haben.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

Das `#box-2`-Element sollte die Eigenschaft `flex-grow` mit dem Wert `2` haben.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
