---
id: 66da326c02141df538f29ba5
title: Costruisci una pagina di checkout
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere un elemento `h1` con il testo `Checkout`.
1. Dovresti avere due elementi `section` immediatamente dopo l'elemento `h1`.
1. Dovresti avere un elemento `h2` con il testo `Your Cart` all'interno della prima sezione.
1. Dovresti avere un'immagine di un elemento nella prima sezione con un testo alternativo appropriato. Puoi usare questa immagine se vuoi: `https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. Dovresti avere un elemento `h2` con il testo `Payment Information` all'interno della seconda sezione.
1. Dovresti avere un elemento `form` all'interno della seconda sezione.
1. Dovresti avere un input con un `id` e `name` di `card-name`, e un `type` di `text` all'interno del modulo e un `label` associato.
1. Dovresti avere un input con un `id` e `name` di `card-number`, e un `type` di `text` all'interno del modulo e un `label` associato.
1. Dovresti avere almeno due elementi `input` con l'attributo `required`.
1. Dovresti includere un elemento `span` con il testo `*` e `aria-hidden` impostato su `true` dentro l'elemento `label` per ogni input obbligatorio, così che i campi obbligatori siano indicati visivamente.
1. Dovresti avere un elemento `p` con un testo di aiuto che spiega il formato richiesto per il numero della carta, posizionato immediatamente dopo l'input del numero della carta. Il `p` dovrebbe avere un `id` di `card-number-help` ed essere referenziato dall'input del numero della carta usando `aria-describedby`.

# --hints--

Dovresti avere un elemento `h1` con il testo `Checkout`.

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

Dovresti avere un solo elemento `h1` nella tua pagina.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

Dovresti avere almeno due elementi `section` dopo l'elemento `h1`.

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

Dovresti avere un elemento `h2` con il testo `Your Cart` all'interno della prima sezione.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

Dovresti avere almeno un'immagine con testo alternativo nella tua prima sezione.

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

Dovresti avere un elemento `h2` con il testo `Payment Information` all'interno della seconda sezione.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

Dovresti avere un elemento `form` all'interno del secondo elemento `section`.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

Dovresti avere un elemento `input` con un `id` e `name` di `card-name` all'interno del modulo.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

Il tuo input per il nome sulla carta dovrebbe avere un `type` di `text`.

```js
const cardNameInput = document.getElementById('card-name');
assert.exists(cardNameInput);
assert.equal(cardNameInput.getAttribute('type'), 'text');
```

Dovresti avere un elemento `input` con un `id` e `name` di `card-number` all'interno del modulo.

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

Il tuo input per il numero della carta dovrebbe avere un `type` di `text`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.exists(cardNumberInput);
assert.equal(cardNumberInput.getAttribute('type'), 'text');
```

Tutti i tuoi elementi `input` che non sono un `type` di `submit` dovrebbero avere un elemento `label` associato.

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
});
```

Dovresti avere almeno due elementi `input` con l'attributo `required`.

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required);
assert.isAtLeast(requiredInputs?.length, 2)
```

Dovresti includere un elemento `span` con il testo `*` dentro l'elemento `label` per ogni input obbligatorio.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.textContent.trim(), '*');
});
```

I tuoi elementi `span` dovrebbero avere `aria-hidden` impostato su `true`.

```js
const requiredInputs = Array.from(document.querySelectorAll('input'))
  .filter(input => input.required);

assert.isNotEmpty(requiredInputs);

requiredInputs.forEach(input => {
  const label = document.querySelector(`label[for="${input.id}"]`) || input.parentElement;
  assert.exists(label);
  assert.equal(label.tagName, "LABEL");
  const span = label?.querySelector('span');
  assert.equal(span?.getAttribute('aria-hidden'), 'true');
});
```

Dovresti avere un elemento `p` con un `id` di `card-number-help` immediatamente dopo l'input del numero della carta.

```js
const cardNumberInput = document.getElementById('card-number');
const cardNumberHelp = document.getElementById('card-number-help');

assert.exists(cardNumberHelp);
const nextElement = cardNumberInput.parentElement?.tagName === 'LABEL'
  ? cardNumberInput.parentElement.nextElementSibling
  : cardNumberInput.nextElementSibling;
assert.equal(nextElement, cardNumberHelp);
```

Il testo di aiuto per il numero della carta non dovrebbe essere vuoto.

```js
const cardNumberHelp = document.getElementById('card-number-help');
assert.isString(cardNumberHelp?.textContent);
assert.isAbove(cardNumberHelp?.textContent.trim().length, 0);
```

Il tuo input per il numero della carta dovrebbe avere `aria-describedby` impostato su `card-number-help`.

```js
const cardNumberInput = document.getElementById('card-number');
assert.equal(cardNumberInput.getAttribute('aria-describedby'), 'card-number-help');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-name" name="card-name" required>
                </div>
                <div>
                    <label for="card-number">Card Number <span aria-hidden="true">*</span></label>
                    <input type="text" id="card-number" name="card-number" aria-describedby="card-number-help" required>
                    <p id="card-number-help">Please enter your 16-digit card number without spaces or dashes.</p>
                </div>
                <div>
                    <label for="expiry-date">Expiry Date <span aria-hidden="true">*</span></label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required>
                </div>
                <div>
                    <label for="cvv">CVV <span aria-hidden="true">*</span></label>
                    <input type="text" id="cvv" name="cvv" required aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
