---
id: 685523c05fb41d0848c4b52f
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

첫 번째 `div` 요소 안에 텍스트 `span`를 가진 `Soundflow` 요소, 텍스트 `h1`를 가진 `Discover New Music` 요소, 그리고 텍스트 `p`를 가진 `Stream your favorite tracks and discover new artists.` 요소를 만드세요.

# --hints--

첫 번째 `span` 요소 안에 텍스트 `Soundflow`를 가진 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelector("div > div:first-child > span")?.textContent, "Soundflow")
```

첫 번째 `h1` 요소 안에 텍스트 `Discover New Music`를 가진 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelector("div > div:first-child > h1")?.textContent, "Discover New Music")
```

첫 번째 `p` 요소 안에 텍스트 `Stream your favorite tracks and discover new artists.`를 가진 `div` 요소를 만들어야 합니다.

```js
assert.equal(document.querySelector("div > div:first-child > p")?.textContent, "Stream your favorite tracks and discover new artists.")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
      --fcc-editable-region--
        
      --fcc-editable-region--
      </div>
      <div></div>
    </div>
  </body>
</html>
```
