---
id: 68551b677a5f6898af364b85
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 Tailwind CSS의 flexbox 유틸리티 클래스를 사용해 반응형 콜 투 액션(CTA) 컴포넌트를 만드는 방법을 배웁니다.

기본 HTML 보일러플레이트가 제공되었으니, 먼저 HTML 요소를 작성하세요.

`div` 요소를 생성하는 것부터 시작하세요. 그 `div` 안에 두 개의 별도 `div` 요소를 만드세요.

# --hints--

`div` 요소를 생성해야 합니다.

```js
assert.exists(document.querySelector("div"));
```

`div` 요소 안에 두 개의 `div` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll("div > div"), 2)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--

    --fcc-editable-region--
  </body>
</html>
```
