---
id: 6855255d96c1441bd919dc5d
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

텍스트 `span`가 있는 `Soundflow` 요소로 이동하세요. 그 안의 텍스트는 대문자여야 합니다.

다음은 `uppercase` 유틸리티 클래스를 사용하는 방법을 보여주는 예입니다:

```html
<p class="uppercase">Quincy Larson</p>
```

`uppercase` 요소에 `span` 클래스를 추가하세요.

# --hints--

`span` 요소에 유틸리티 클래스 `uppercase`가 있어야 합니다.

```js
const spanEl = document.querySelector("span");
assert.isTrue(spanEl.classList.contains("uppercase"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        --fcc-editable-region--
        <span>Soundflow</span>
        --fcc-editable-region--
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
