---
id: 68552464ef378e1051a55c28
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

두 번째 `div` 안에 앵커 요소 두 개를 만드세요. 첫 번째 앵커 요소에는 텍스트 `Learn more`를, 두 번째 앵커 요소에는 텍스트 `Start listening`를 넣으세요.

두 앵커 요소의 `href` 속성은 둘 다 `#`로 설정하세요.

# --hints--

두 번째 `a` 요소 안에 `div` 요소 두 개를 만드세요.

```js
assert.lengthOf(document.querySelectorAll("div > div:nth-child(2) > a"), 2)
```

첫 번째 `a` 요소에는 텍스트 `Learn more`와 두 번째 `href` 요소 안에 `#` 속성으로 `div`를 넣으세요.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[0]?.textContent, "Learn more")
assert.equal(links[0]?.getAttribute("href"), "#")
```

두 번째 `a` 요소에는 텍스트 `Start listening`와 두 번째 `href` 요소 안에 `#` 속성으로 `div`를 넣으세요.

```js
const links = document.querySelectorAll("div > div:nth-child(2) > a")

assert.equal(links[1]?.textContent, "Start listening")
assert.equal(links[1]?.getAttribute("href"), "#")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
