---
id: 6858ebf2552046ab3de9f17c
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

이 시점에서 `div`의 콘텐츠 오른쪽에 너무 많은 공간이 있는 것을 알 수 있습니다. 너비를 추가하여 이 공간을 최소화할 수 있습니다.

중간 및 그 이상의 화면 크기, 즉 768픽셀 이상인 화면에서 너비를 50%로 줄이려면 `md:w-1/2`에 `div` 클래스를 추가하세요.

이 효과를 확인하려면 미리보기 탭을 확장하세요.

# --hints--

외부 `div` 요소에 유틸리티 클래스 `md:w-1/2`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("md:w-1/2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
