---
id: 6858f90c54647a48ddb8822d
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

이 시점에서, 즉 바깥쪽 `div` 안에 있는 두 개의 `div` 요소는 화면 크기에 상관없이 서로 겹쳐 쌓일 것입니다.

그 이유는 감싸는 `flex-col` 요소에 `div` 클래스가 있기 때문입니다. 이를 고치려면 `lg:` 접두사를 사용해 `flex-row`에 `div` 클래스를 추가하세요.

즉, 화면이 큰 기기에서는 flex 방향이 `row`가 됩니다. 다시 효과를 확인하려면 브라우저 너비를 더 크게 조절하세요.

# --hints--

바깥쪽 `div` 요소에 유틸리티 클래스 `lg:flex-row`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("lg:flex-row"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
