---
id: 6858eadd3144c89d45e02bc7
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

보면, 바깥쪽 `div`가 위쪽에 너무 붙어 있고 그 안의 내용이 왼쪽으로 너무 몰려 있습니다. 그래서 이제 간격을 조절해야 합니다.

`mt-8`에 `div` 클래스를 추가해 약간 아래로 내리고, `p-4` 클래스를 추가해 텍스트가 왼쪽에서 떨어지도록 하세요.

# --hints--

바깥쪽 `div` 요소에 유틸리티 클래스 `mt-8`를 추가해야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mt-8"));
```

바깥쪽 `div` 요소에 유틸리티 클래스 `p-4`를 추가해야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("p-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
