---
id: 6858f3a67aa8580ff37016c7
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

미리보기 탭의 너비를 모바일 기기와 동일한 너비로 다시 조정하면 콘텐츠가 항상 나란히 배치된 것을 볼 수 있습니다. 모바일 기기에서는 이렇게 되면 안 됩니다.

이를 고치려면 `flex-col`에 `div` 클래스를 추가하세요. 그러면 flex 방향이 열 방향으로 바뀌어 요소들이 위아래로 쌓이게 됩니다.

`flex-col` 유틸리티 클래스를 사용하는 방법은 다음과 같습니다.

```html
<div class="flex flex-col">
  <element></element>
  <element></element>
</div>
```

# --hints--

열리는 `div` 요소에 유틸리티 클래스 `flex-col`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex-col"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
