---
id: 6858e9f87e4faa903267d810
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

외부 `div`의 텍스트가 기본 검정색이라 보기 쉽지 않습니다. 그래서 `text-white`에 `div` 클래스를 추가하세요.

# --hints--

외부 `div` 요소에 유틸리티 클래스 `text-white`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
