---
id: 6858e6408a0999603b2a4fee
title: 14단계
challengeType: 0
dashedName: step-14
---

# --description--

`h1`에 `font-bold`과 `text-4xl` 클래스를 추가하세요. 이렇게 하면 텍스트가 굵고 커지며 줄 간격도 넓어집니다.

또한 `my-4`의 위아래 여백을 16픽셀, 즉 `h1`만큼 늘리기 위해 `1rem` 클래스를 추가하세요.

# --hints--

`h1` 요소에 유틸리티 클래스 `font-bold`가 있어야 합니다.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("font-bold"));
```

`h1` 요소에 유틸리티 클래스 `text-4xl`가 있어야 합니다.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("text-4xl"));
```

`h1` 요소에 유틸리티 클래스 `my-4`가 있어야 합니다.

```js
const h1El = document.querySelector("h1");
assert.isTrue(h1El.classList.contains("my-4"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
         --fcc-editable-region--
        <h1>Discover New Music</h1>
         --fcc-editable-region--
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
  </body>
</html>
```
