---
id: 6858ee73ed8106ccb214942a
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

모든 내용이 왼쪽에 붙었으니, x축에 자동 여백을 추가해 전체 내용을 가운데로 맞추세요.

그에 해당하는 Tailwind 클래스를 외부 `div` 요소에 추가하세요. 효과를 확인하려면 미리보기 탭을 조정해 전체 내용이 가운데 정렬되는지 확인하세요.

# --hints--

외부 `div` 요소에 유틸리티 클래스 `mx-auto`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("mx-auto"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--    
  </body>
</html>
```
