---
id: 6858e7e23bf9e475e14cb931
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

HTML 구조를 완성했으니 이제 Tailwind로 적절한 스타일링을 적용할 차례입니다.

전체 내용을 감싸는 `div` 요소부터 시작하세요. `class`에 `bg-indigo-600` 값을 주어 배경색을 설정하세요.

# --hints--

외부 `div`에는 `class` 속성을 `bg-indigo-600`로 설정해야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("bg-indigo-600"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    --fcc-editable-region--
    <div>
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
   --fcc-editable-region--
  </body>
</html>
```
