---
id: 6858ef974c8ad5dcd7abb9f1
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

`div` 내부의 내용은 더 나은 간격과 배열이 필요합니다. 나머지 간격과 배열을 위해서는 flexbox를 사용해야 합니다.

두 개의 자식 `flex` 요소가 나란히 앉도록 `div`에 `div` 클래스를 추가하세요.

또한, 내용 주위와 사이의 공간이 고르게 분배되도록 `justify-around`를 추가하고, 전체 내용이 `items-center` 안에서 최종적으로 가운데 정렬되도록 `div`를 추가하세요.

# --hints--

외부 `div` 요소에 유틸리티 클래스 `flex`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("flex"));
```

외부 `div` 요소에 유틸리티 클래스 `justify-around`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("justify-around"));
```

외부 `div` 요소에 유틸리티 클래스 `items-center`가 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("items-center"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
       <div>
         <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
