---
id: 685975c2505e3db255720561
title: 18단계
challengeType: 0
dashedName: step-18
---

# --description--

마지막으로, 버튼들 사이와 위쪽 간격을 더 넓게 조정할 수 있습니다. 그렇게 하려면 두 버튼을 감싸는 `div` 요소를 선택해야 합니다. 그래서 그 `flex`에 `gap-2`, `mt-6`, `div` 클래스를 추가하세요.

이렇게 하면 CTA 컴포넌트 워크숍이 완료됩니다!

# --hints--

앵커 요소들을 감싸는 `div` 요소에 유틸리티 클래스 `flex`를 추가하세요.

```js
const targetDiv = document.querySelectorAll("div")[2];
assert.isTrue(targetDiv.classList.contains("flex"));
```

앵커 요소들을 감싸는 `div` 요소에 유틸리티 클래스 `gap-2`를 추가하세요.

```js
const targetDiv = document.querySelectorAll("div")[2];
assert.isTrue(targetDiv.classList.contains("gap-2"));
```

앵커 요소들을 감싸는 `div` 요소에 유틸리티 클래스 `mt-6`를 추가하세요.

```js
const targetDiv = document.querySelectorAll("div")[2];
assert.isTrue(targetDiv.classList.contains("mt-6"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
--fcc-editable-region--
      <div >
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600 px-4 py-2 font-semibold rounded"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2 font-semibold rounded"
         >
          Start listening
        </a>
      </div>
--fcc-editable-region--
    </div>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div class="flex gap-2 mt-6">
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600 px-4 py-2 font-semibold rounded"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2 font-semibold rounded"
         >
          Start listening
        </a>
      </div>
    </div>
  </body>
</html>
```
