---
id: 685972d14c6d8182fbae71f1
title: 16단계
challengeType: 0
dashedName: step-16
---

# --description--

앵커 태그를 버튼처럼 보이게 하려면 각 태그에 클래스 `px-4`와 `py-2`을 추가하세요. `px-4`는 좌우 패딩을 `1rem`만큼 늘려주고, `py-2`는 상하 패딩을 `0.5rem`만큼 늘려줍니다.

# --hints--

첫 번째 앵커 요소에는 유틸리티 클래스 `px-4`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("px-4"));
```

첫 번째 앵커 요소에는 유틸리티 클래스 `py-2`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("py-2"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `px-4`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("px-4"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `py-2`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("py-2"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white"
         >
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
