---
id: 68597114730d7c6b8838d204
title: 15단계
challengeType: 0
dashedName: step-15
---

# --description--

`h1` 아래 단락은 잘 되어 있으니 추가 스타일링이 필요하지 않습니다. 따라서 두 번째 `div` 요소 안에 있는 앵커 요소로 넘어가세요.

첫 번째 앵커 요소에 `bg-white`, `hover:bg-gray-100`, `text-indigo-600` 클래스를 주세요. 이렇게 하면 흰색 배경, 인디고 텍스트 색상, 그리고 마우스를 올렸을 때 약간 밝은 배경이 적용됩니다.

두 번째 앵커 요소에는 `bg-pink-500`, `hover:bg-pink-600`, `text-white` 클래스를 주세요.

# --hints--

첫 번째 앵커 요소에는 유틸리티 클래스 `bg-white`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("bg-white"));
```

첫 번째 앵커 요소에는 유틸리티 클래스 `hover:bg-gray-100`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("hover:bg-gray-100"));
```

첫 번째 앵커 요소에는 유틸리티 클래스 `text-indigo-600`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("text-indigo-600"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `bg-pink-500`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("bg-pink-500"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `hover:bg-pink-600`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("hover:bg-pink-600"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `text-white`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("text-white"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        --fcc-editable-region--
        <a href="#">
          Learn more
        </a>
        <a href="#">
          Start listening
        </a>
        --fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
