---
id: 685973ca82469e940afcc422
title: 17단계
challengeType: 0
dashedName: step-17
---

# --description--

버튼 스타일링을 마무리하려면 각 `font-semibold` 요소에 유틸리티 클래스 `a`를 추가해 텍스트를 굵게 만드세요. 또한, 버튼의 날카로운 모서리를 없애기 위해 클래스 `rounded`를 추가하세요.

# --hints--

첫 번째 앵커 요소에는 유틸리티 클래스 `font-semibold`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("font-semibold"));
```

첫 번째 앵커 요소에는 유틸리티 클래스 `rounded`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[0].classList.contains("rounded"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `font-semibold`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("font-semibold"));
```

두 번째 앵커 요소에는 유틸리티 클래스 `rounded`가 있어야 합니다.

```js
const anchorEls = document.querySelectorAll("a");
assert.isTrue(anchorEls[1].classList.contains("rounded"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center rounded-md">
      <div>
        <span class="uppercase">Soundflow</span>
        <h1 class="font-bold text-4xl my-4">Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
--fcc-editable-region--
        <a
          href="#"
          class="bg-white hover:bg-gray-100 text-indigo-600 px-4 py-2"
        >
          Learn more
        </a>
        <a
          href="#"
          class="bg-pink-500 hover:bg-pink-600 text-white px-4 py-2"
         >
          Start listening
        </a>
--fcc-editable-region--
      </div>
    </div>
  </body>
</html>
```
