---
id: 68591755e8cd86d27699d282
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

마지막으로, 바깥쪽 `div` 요소의 날카로운 모서리를 없애려면 클래스 `rounded-md`을 추가하세요.

# --hints--

바깥쪽 `div` 요소에 유틸리티 클래스 `rounded-md`이 있어야 합니다.

```js
const firstDiv = document.querySelectorAll("div")[0];
assert.isTrue(firstDiv.classList.contains("rounded-md"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>CTA component</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
--fcc-editable-region--
    <div class="bg-indigo-600 text-white mt-8 p-4 md:w-1/2 mx-auto flex flex-col lg:flex-row justify-around items-center">
      <div>
        <span>Soundflow</span>
        <h1>Discover New Music</h1>
        <p>Stream your favorite tracks and discover new artists.</p>
      </div>
      <div>
        <a href="#">Learn more</a>
        <a href="#">Start listening</a>
      </div>
    </div>
--fcc-editable-region--
  </body>
</html>
```
