---
id: 6690be42942f421648b119eb
title: Hatua ya 7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

Ongeza kipengele cha `img` chini ya kipengele cha `p`. Katika hatua hii, hakuna picha itaonyeshwa kwenye kivinjari.

# --hints--

Kipengele chako cha `img` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element should not have a closing tag `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Unapaswa kuwa na kipengele kimoja tu cha `img`. Ondoa ziada yoyote.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Kipengele chako cha `img` kinapaswa kuwa chini ya kipengele cha `p`. Unazo katika mpangilio mbaya.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


