---
id: 6698429ba5761b39f8aad814
title: Hatua ya 17
challengeType: 0
dashedName: step-17
---

# --description--

It is time to add a new section. Add a second `section` element below the existing `section` element.

# --hints--

Kipengee chako cha `section` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Unapaswa kuongeza tu tagi moja ya ufunguzi ya `section`. Tafadhali ondoa ziada yoyote.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Kipengee chako cha `section` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g).length, 2);
```

Unapaswa kuongeza tu tagi moja ya kufunga ya `section`. Tafadhali ondoa ziada yoyote.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Kipengele cha pili cha `section` hakipaswi kuwekwa ndani ya kipengee cha kwanza cha `section`.

```js
const childrenOf1stSection = [
  ...document.querySelector('main > section').children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Vipengele vyote viwili vya `section` vinapaswa kuwa kati ya tagi za kufungua na kufunga za kipengele cha `main`.

```js
const childrenOfMain = [...document.querySelector('main').children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```


