---
id: 669842cee55b683aa13ee758
title: Hatua ya 18
challengeType: 0
dashedName: step-18
---

# --description--

Within the second `section` element, add a new `h2` element with the text `Cat Lists`.

# --hints--

Kipengee chako cha `section` kinapaswa kuwa na tagi ya ufunguzi. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
assert.lengthOf(code.match(/<\/section>/g), 2);
```

Kipengele chako cha `h2` kinafaa kuwa na tagi ya kufungua. Tagi za ufunguzi zina sintaksia hii: `<elementName>`.

```js
assert.lengthOf(document.querySelectorAll('h2'), 2);
```

Kipengele chako cha `h2` kinapaswa kuwa na tagi ya kufunga. Tagi za kufunga zina `/` mara baada ya herufi ya `<`.

```js
assert.lengthOf(code.match(/<\/h2\>/g), 2);
```

Kipengele chako cha pili cha `h2` kinafaa kuwa juu ya tagi ya kufunga ya kipengele cha pili cha `section`. Haiko katika nafasi sahihi.

```js
const secondSection = document.querySelectorAll('section')[1];
assert.equal(secondSection.lastElementChild.nodeName, 'H2');
```

Kipengele chako cha pili cha `h2` kinapaswa kuwa na maandishi `Cat Lists`. Aidha umesahau maandishi, au kuna makosa ya maandishi.

```js
assert.equal(
  document
    .querySelectorAll('main > section')[1]
    .lastElementChild.innerText.toLowerCase(), 'cat lists'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
--fcc-editable-region--
      <section>

      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```


