---
id: 6ac09f8420b4b67e075d2cd2
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop erstellen Sie Ihr erstes Python-Programm.

Dies ist Ihr Code-Editor, in dem Sie Python-Code schreiben. Zeile 1 ist markiert. Gehen Sie dorthin und geben Sie Folgendes ein:

```py
'Hello'
```

Achten Sie darauf, am Anfang der Zeile zu beginnen, und vergessen Sie die Anführungszeichen nicht.

Wenn Sie fertig sind, klicken Sie auf die Schaltfläche „Prüfen Sie Ihren Code“, um zu sehen, ob er korrekt ist.

# --hints--

Sie sollten `'Hello'` eingeben. Achten Sie darauf, am Anfang der Zeile zu beginnen, und vergessen Sie die Anführungszeichen nicht.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""(['"])\\s*hello\\s*\\1"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
