---
id: 6ac0aaba914d81b32567c076
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

`print` ist eine eingebaute Funktion, die Sie verwenden können, um Ausgabe an das Terminal zu senden. Wie Sie sehen, ist der Text `Hello` jetzt im Terminal sichtbar.

Geben Sie nun den Text `freeCodeCamp` im Terminal aus, indem Sie zur markierten Zeile gehen und Folgendes eingeben:

```py
print('freeCodeCamp')
```

Achten Sie darauf, am Anfang der Zeile zu beginnen.

# --hints--

Ihr Code sollte `print('freeCodeCamp')` enthalten. Achten Sie darauf, am Anfang der Zeile zu beginnen.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*freecodecamp\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
--fcc-editable-region--

--fcc-editable-region--
```
