---
id: 6ac0a4463662d9738c393e21
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Im vorherigen Schritt haben Sie einen **String** erstellt: eine Folge von Zeichen, die entweder in einfache oder doppelte Anführungszeichen eingeschlossen sind. In der Programmierung werden Strings verwendet, um Text darzustellen.

Das **Terminal** ist der Ort, an dem Sie die Ausgabe Ihres Codes sehen können.

Um eine Ausgabe sehen zu können, ändern Sie Ihre bestehende Codezeile in:

```py
print('Hello')
```

Achten Sie darauf, am Anfang der Zeile zu beginnen.

# --hints--

Ihr Code sollte `print('Hello')` enthalten. Achten Sie darauf, am Anfang der Zeile zu beginnen.

```js
({ test: () => runPython(`
import re
stmts = _Node(_code)
pattern = r"""print\\s*\\((['"])\\s*hello\\s*\\1\\s*\\)"""
assert any(
  re.fullmatch(pattern, str(stmt).strip(), flags=re.IGNORECASE)
  for stmt in stmts
)`) });
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
'Hello'
--fcc-editable-region--
```
