---
id: 6ac0b35f3b29a27f7e207295
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Im vorherigen Schritt haben Sie die Variable `sentence` erstellt. Geben Sie nun ihren Wert im Terminal aus, indem Sie `print(sentence)` in die markierte Zeile eingeben.

Damit haben Sie Ihr erstes Python-Programm abgeschlossen.

# --hints--

Ihr Code sollte `print(sentence)` enthalten. Achten Sie darauf, am Anfang der Zeile zu beginnen.

```js
({ test: () => assert(runPython(`_Node(_code).has_stmt("print(sentence)")`)) });
```

# --seed--

## --seed-contents--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```py
print('Hello')
print('freeCodeCamp')

sentence = 'I am learning Python'
print(sentence)
```
