---
id: 68f890b73bbfc9efa7379e7b
title: 在有錯誤的應用程式中偵錯型別強制轉換錯誤
challengeType: 26
dashedName: debug-type-coercion-errors-in-a-buggy-app
---

# --description--

你剛加入一家區域的網頁開發工作室，你的第一個指派是清理前一批開發者留下來的一些有錯誤的程式碼。

他們嘗試了簡單的算術操作，但出了問題，結果沒有意義。你的任務是審查、偵錯並修正程式碼，使其如預期般運作且更易於閱讀。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。 

**使用者故事：**

1. 你應該有一個名為 `firstResult` 的變數，正確地將數字 `5` 和 `10` 相加以產生值 `15`。
2. 你應該有一個名為 `secondResult` 的變數，正確地從 `8` 中減去 `5`，以產生值 `3`。
3. 你應該有一個名為 `thirdResult` 的變數，透過將兩個數字相加來產生值 `6`。將目前在表達式中使用的布林值取代為數字。
4. 你應該有一個名為 `fourthResult` 的變數，透過將兩個數字相加來產生值 `8`。將目前在表達式中使用的布林值取代為數字。
5. 你應該有一個名為 `fifthResult` 的變數，正確地將兩個數字相乘以取得乘積 `20`。將目前在表達式中使用的字串取代為數字。
6. 你應該有一個名為 `sixthResult` 的變數，正確地將兩個數字相加以產生值 `22`。將目前在表達式中使用的 `null` 取代為一個數字。

# --hints--

你應該有一個名為 `firstResult` 的變數。

```js
assert.match(code, /const\s+firstResult\s*=/);
```

`firstResult` 應該具有數值 `15`。

```js
assert.strictEqual(firstResult, 15);
```

你應該有一個名為 `secondResult` 的變數。

```js
assert.match(code, /const\s+secondResult\s*=/);
```

`secondResult` 應該具有數值 `3`。

```js
assert.strictEqual(secondResult, 3);
```

你應該有一個名為 `thirdResult` 的變數。

```js
assert.match(code, /const\s+thirdResult\s*=/);
```

`thirdResult` 應該具有數值 `6`。

```js
assert.strictEqual(thirdResult, 6);
```

`thirdResult` 不應包含布林值（`true` 或 `false`）。

```js
const match = code.match(/const\s+thirdResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

你應該有一個名為 `fourthResult` 的變數。

```js
assert.match(code, /const\s+fourthResult\s*=/);
```

`fourthResult` 應該具有數值 `8`。

```js
assert.strictEqual(fourthResult, 8);
```

`fourthResult` 不應包含布林值（`true` 或 `false`）。

```js
const match = code.match(/const\s+fourthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\btrue\b|\bfalse\b/);
```

你應該有一個名為 `fifthResult` 的變數。

```js
assert.match(code, /const\s+fifthResult\s*=/);
```

`fifthResult` 應該具有數值 `20`。

```js
assert.strictEqual(fifthResult, 20);
```

`fifthResult` 不應包含字串乘法（例如 "10" * 2）。

```js
const match = code.match(/const\s+fifthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /["'`].*["'`]\s*\*/);
```

你應該有一個名為 `sixthResult` 的變數。

```js
assert.match(code, /const\s+sixthResult\s*=/);
```

`sixthResult` 應該有數值 `22`。

```js
assert.strictEqual(sixthResult, 22);
```

`sixthResult` 的表達式中不應包含 `null`。

```js
const match = code.match(/const\s+sixthResult\s*=\s*([^;]+)/);
assert.isNotNull(match);
const expression = match[1];
assert.notMatch(expression, /\bnull\b/);
```

# --seed--

## --seed-contents--

```js
const firstResult = 5 + "10";
console.log(`5 + 10 = ${firstResult}`);

const secondResult = "Eight" - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = true + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = false + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = "10" * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = null + 22;
console.log(`0 + 22 = ${sixthResult}`);
```

# --solutions--

```js
const firstResult = 5 + 10;
console.log(`5 + 10 = ${firstResult}`);

const secondResult = 8 - 5;
console.log(`8 - 5 = ${secondResult}`);

const thirdResult = 1 + 5;
console.log(`1 + 5 = ${thirdResult}`);

const fourthResult = 0 + 8;
console.log(`0 + 8 = ${fourthResult}`);

const fifthResult = 10 * 2;
console.log(`10 * 2 = ${fifthResult}`);

const sixthResult = 0 + 22;
console.log(`0 + 22 = ${sixthResult}`);
```
