---
id: 587d7dbc367417b2b2512bb1
title: 創建一個簡單的 JSX 元素
challengeType: 6
forumTopicId: 301390
dashedName: create-a-simple-jsx-element
---

# --description--

React 是由 Facebook 創建並維護的開放原始碼視圖函式庫。它是呈現現代網頁應用程式使用者介面的絕佳工具。

React 使用一種稱為 JSX 的 JavaScript 語法擴充套件，讓你可以直接在 JavaScript 中撰寫 HTML。這有幾個利益。它讓你能在 HTML 中使用 JavaScript 的完整程式能力，並有助於保持你的程式碼可讀性。在大多數情況下，JSX 與你已經學過的 HTML 類似，但有一些關鍵差異會在這些挑戰中涵蓋。

例如，因為 JSX 是 JavaScript 的語法擴充套件，你實際上可以直接在 JSX 中撰寫 JavaScript。要做到這一點，你只需將你想當作 JavaScript 處理的程式碼包含在大括號中：`{ 'this is treated as JavaScript code' }`。請記住這點，因為它會在未來的幾個挑戰中使用。

但是，因為 JSX 並非有效的 JavaScript，JSX 程式碼必須被編譯成 JavaScript。轉譯器 Babel 是這個處理程序中常用的工具。為了方便起見，它已經在這些挑戰的幕後自動添加。如果你寫出語法無效的 JSX，你會看到這些挑戰中的第一個測試失敗。

值得注意的是，在底層挑戰是呼叫 `ReactDOM.render(JSX, document.getElementById('root'))`。這個函式呼叫會將你的 JSX 放入 React 自己輕量級的 DOM 表現中。React 接著使用其自身 DOM 的螢幕快照來最佳化只更新實際 DOM 的特定零件。

# --instructions--

目前的程式碼使用 JSX 將一個 `div` 元素指定給常數 `JSX`。將 `div` 取代為 `h1` 元素，並在其中添加文字 `Hello JSX!`。

# --hints--

常數 `JSX` 應該傳回一個 `h1` 元素。

```js
assert(JSX.type === 'h1');
```

`h1` 頁籤應該包含文字 `Hello JSX!`

```js
assert(Enzyme.shallow(JSX).contains('Hello JSX!'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(JSX, document.getElementById('root'))
```

## --seed-contents--

```jsx
const JSX = <div></div>;
```

# --solutions--

```jsx
const JSX = <h1>Hello JSX!</h1>;
```
