---
id: 5a24c314108439a4d4036176
title: 使用 State 切換元素
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

有時候你可能需要在更新狀態時知道先前的狀態。然而，狀態更新可能是非同步的－這表示 React 可能會將多個 `setState()` 呼叫批次成單一更新。這表示你無法在計算下一個值時依賴 `this.state` 或 `this.props` 的先前值。因此，你不應該使用像這樣的程式碼：

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

相反地，你應該傳遞一個函式（程式）給 `setState`，讓你可以存取 state 和 props。使用帶有 `setState` 的函式（程式）可保證你正在使用最新的 state 和 props 值。這表示上述內容應該重寫為：

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

如果你只需要 `state`，也可以使用沒有 `props` 的表單：

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

請注意，你必須將物件字面值包裹在小括弧中，否則 JavaScript 會認為它是一個程式碼區塊。

# --instructions--

`MyComponent` 有一個 `visibility` 屬性，初始值為 `false`。當 `visibility` 的值為 true 時，呈現方法會傳回一個視圖；當其為 false 時，則傳回另一個視圖。

目前，沒有辦法更新組件的 `state` 中的 `visibility` 屬性。該值應該在 true 和 false 之間切換。按鈕上有一個點擊處理常式，會觸發一個名為 `toggleVisibility()` 的類別方法。將一個函式傳遞給 `setState` 以定義此方法，使得當該方法被呼叫時，`visibility` 的 `state` 會切換到相反的值。如果 `visibility` 是 `false`，該方法會將其設為 `true`，反之亦然。

最後，點擊按鈕以查看基於其 `state` 的組件條件呈現。

**提示：** 別忘了在 `constructor` 中將 `this` 關鍵字綁定到該方法！

# --hints--

`MyComponent` 應該傳回一個包含 `button` 的 `div` 元素。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

`MyComponent` 的狀態應該初始化，並將 `visibility` 屬性設為 `false`。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

點擊按鈕元素應該在狀態中切換 `visibility` 屬性於 `true` 和 `false` 之間，並有條件地呈現 `h1` 元素。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

應該將匿名函式（程式）傳遞給 `setState`。

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

不應在 `setState` 內使用 `this`

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
