---
id: 5a24c314108439a4d4036180
title: 使用 shouldComponentUpdate 來最佳化重新呈現
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

到目前為止，如果任何組件接收到新的 `state` 或新的 `props`，它會重新呈現自己和所有子組件。這通常沒問題。但 React 提供了一個生命週期方法，當子組件接收到新的 `state` 或 `props` 時，你可以呼叫它，並明確宣告組件是否應該更新。這個方法是 `shouldComponentUpdate()`，它接受 `nextProps` 和 `nextState` 作為參數。

這個方法是最佳化效率的有用方式。例如，預設行為是當你的組件接收到新的 `props` 時會重新呈現，即使 `props` 沒有改變。你可以使用 `shouldComponentUpdate()` 透過比較 `props` 來防止這種情況。該方法必須傳回一個 `布林值`，告訴 React 是否要更新組件。你可以比較目前的 `props`（`this.props`）和下一個 `props`（`nextProps`）來判斷是否需要更新，並相應地傳回 `true` 或 `false`。

# --instructions--

在名為 `OnlyEvens` 的組件中添加了 `shouldComponentUpdate()` 方法。目前，此方法傳回 `true`，因此每次接收到新的 `props` 時，`OnlyEvens` 都會重新呈現。修改此方法，使 `OnlyEvens` 僅在其新 props 的 `value` 為偶數時才更新。點擊 `Add` 按鈕，並在瀏覽器的主控台中觀察生命週期掛鉤被觸發時事件的順序。

# --hints--

`Controller` 組件應該將 `OnlyEvens` 組件作為子組件呈現。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

`shouldComponentUpdate` 方法應該定義在 `OnlyEvens` 組件上。

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

`OnlyEvens` 組件應該傳回一個 `h1` 頁籤，該頁籤呈現 `this.props.value` 的值。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` 應該只在 `nextProps.value` 為偶數時重新呈現。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Controller />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
