---
id: 5a24c314108439a4d4036189
title: 根據組件狀態有條件地更改行內 CSS
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

到目前為止，你已經看到多個條件呈現和行內樣式的應用。這裡有另一個結合這兩個主題的範例。你也可以根據 React 組件的狀態有條件地呈現 CSS。為此，你需要檢查一個條件，如果該條件成立，就修改指定給呈現方法中 JSX 元素的 styles 物件。

這種範式很重要，因為它與更傳統的透過直接修改 DOM 元素來套用樣式的方法有劇烈的轉變（這在 jQuery 中非常常見，例如）。在那種方法中，你必須追蹤元素何時改變，並且直接控制代碼實際的操作。追蹤變化可能變得困難，可能會使你的使用者介面變得不可預測。當你根據條件設定樣式物件時，你描述的是使用者介面應該如何根據應用程式的狀態呈現。資訊流動是單向且明確的。這是在使用 React 撰寫應用程式時偏好的方法。

# --instructions--

程式碼編輯器有一個帶有樣式化邊框的簡單受控輸入組件。如果使用者在輸入框中輸入超過 15 個字元的文字，你想將此邊框樣式設為紅色。為此添加一個條件來檢查，並且如果條件成立，將輸入邊框樣式設為 `3px solid red`。你可以透過在輸入框中輸入文字來試用。

# --hints--

`GateKeeper` 組件應該呈現一個 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`GateKeeper` 組件應該以狀態鍵 `input` 初始化，並設為空字串。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

`GateKeeper` 組件應該呈現一個 `h3` 頁籤和一個 `input` 頁籤。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

`input` 頁籤的 `border` 屬性應該最初設定為 `1px solid black` 的樣式。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

如果狀態中的輸入值長度超過 15 個字元，則應該為 `input` 頁籤套用 `3px solid red` 的邊框樣式。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<GateKeeper />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
