---
id: 5a24c314108439a4d4036173
title: 使用 this.setState 設定狀態
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

先前的挑戰涵蓋了組件的 `state` 以及如何在 `constructor` 中初始化 state。也有一種方法可以改變組件的 `state`。React 提供了一個用於更新組件 `state` 的方法，稱為 `setState`。你可以在你的組件類別中這樣呼叫 `setState` 方法：`this.setState()`，並傳入一個包含鍵值對的物件。鍵是你的 state 屬性，值是更新後的 state 資料。例如，如果我們在 state 中儲存了一個 `username` 並想要更新它，會是這樣：

```jsx
this.setState({
  username: 'Lewis'
});
```

React 預期你絕不會直接修改 `state`，而是在狀態改變時總是使用 `this.setState()`。此外，你應該注意到 React 可能會將多個狀態更新合併成批次以提升效率。這表示透過 `setState` 方法的狀態更新可能是非同步的。`setState` 方法有一種替代語法，可以解決這個問題。這種情況很少需要，但記住它是很好的！請參考我們的 <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">React 文章</a> 以獲得更多細節。

# --instructions--

在程式碼編輯器中有一個 `button` 元素，其具有一個 `onClick()` 處理常式。當瀏覽器中的 `button` 接收到點擊事件時，會觸發此處理常式，並執行定義在 `MyComponent` 上的 `handleClick` 方法。在 `handleClick` 方法中，使用 `this.setState()` 更新組件的 `state`。將 `state` 中的 `name` 屬性設定為相等於字串 `React Rocks!`。

點擊按鈕並觀察呈現狀態的更新。如果你目前還不完全了解點擊處理常式的程式碼如何運作，也不用擔心。這部分會在接下來的挑戰中涵蓋。

# --hints--

`MyComponent` 的狀態應該以鍵值對 `{ name: Initial State }` 初始化。

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` 應該呈現一個 `h1` 標題元素。

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

算繪後的 `h1` 標題組件元素應包含從組件狀態算繪的文字。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

在 `MyComponent` 上呼叫 `handleClick` 方法應該會將狀態中的 name 屬性設定為相等於 `React Rocks!`。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
