---
id: 5a24c314108439a4d4036181
title: 介紹行內樣式
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

還有其他複雜的概念，為你的 React 程式碼添加強大的功能。但你可能會想知道如何為你在 React 中創建的那些 JSX 元素設計樣式這個較簡單的問題。你很可能知道，這不會完全像操作 HTML 一樣，因為<a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">你套用類別到 JSX 元素的方式</a>。

如果你從樣式表匯入樣式，基本上沒有太大差別。你使用 `className` 屬性將類別套用到你的 JSX 元素，並在樣式表中將樣式套用到該類別。另一個選項是套用行內樣式，這在 ReactJS 開發中非常常見。

你對 JSX 元素套用行內樣式的方式類似於在 HTML 中的做法，但有一些 JSX 的差異。以下是一個 HTML 中行內樣式的範例：

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

JSX 元素使用 `style` 屬性，但因為 JSX 的轉譯方式，你不能將值設定為 `string`。相反地，你必須將它設定為 JavaScript 的 `object`。以下是一個範例：

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

注意我們如何將 `fontSize` 屬性寫成 camelCase？這是因為 React 不接受 style 物件中的 kebab-case 鍵。React 會在 HTML 中為我們套用正確的屬性名稱。

# --instructions--

為你的 `div` 添加一個 `style` 屬性到程式碼編輯器中，將文字顏色設為紅色，字型大小設為 `72px`。

請注意，你可以選用將字型大小設定為數字，省略單位 `px`，或寫成 `"72px"`。
# --hints--

該組件應該呈現一個 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`div` 元素的顏色應該是 `red`。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

`div` 元素應該有 `72px` 的字型大小。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Colorful />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
