---
id: 5a24c314108439a4d403616d
title: 使用 PropTypes 定義你期望的 Props
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React 提供有用的型別檢查特性，以驗證組件接收到的 props 是否為正確的型別。 例如，你的應用程式進行 API 呼叫以取得你預期為陣列的資料，然後將該資料作為 prop 傳遞給組件。你可以在組件上設定 `propTypes`，以要求資料必須為 `array` 型別。當資料為其他任何型別時，這將發出有用的警告。

當你事先知道一個 prop 的型別時，設定 `propTypes` 被視為最佳實務。你可以用定義 `defaultProps` 的相同方式，為組件定義一個 `propTypes` 屬性。這樣做會檢查給定鍵的 props 是否存在且具有指定的型別。以下是一個範例，要求名為 `handleClick` 的 prop 必須是 `function` 型別：

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

在上述範例中，`PropTypes.func` 部分會檢查 `handleClick` 是否為函式。為你的組件添加 `isRequired` 告訴 React `handleClick` 是必填的屬性。如果沒有提供該屬性，你將看到警告。也請注意 `func` 表示 `function`。在七種 JavaScript 原生型別中，`function` 和 `boolean`（寫作 `bool`）是唯一兩個使用不尋常拼寫的型別。除了原生型別之外，還有其他可用的型別。例如，你可以檢查屬性是否為 React 元素。請參考 <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">文件</a> 以了解所有選項。

**注意：** 從 React v15.5.0 起，`PropTypes` 是獨立於 React 進行匯入的，方式如下：`import PropTypes from 'prop-types';`

# --instructions--

為 `Items` 組件定義 `propTypes`，以要求 `quantity` 作為屬性並驗證其型別為 `number`。

# --hints--

`ShoppingCart` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`Items` 組件應該呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

`Items` 組件應該包含一個 `propTypes` 檢查，以要求 `quantity` 的值並確保其值為數字。

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --seed--

## --before-user-code--

```jsx
var PropTypes = {
  number: { isRequired: true }
};
```

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
