---
id: 5a24c314108439a4d4036164
title: 使用複合創建組件
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

現在我們將探討如何將多個 React 組件組合在一起。想像你正在建置一個應用程式，並且已經創建了三個組件：`Navbar`、`Dashboard` 和 `Footer`。

要將這些組件組合在一起，你可以創建一個 `App` *父* 組件，該組件將這三個組件中的每一個作為 *子* 組件來呈現。要在 React 組件中將組件呈現為子組件，你需要在 JSX 中以訂製的 HTML 頁籤形式包含該組件名稱。舉例來說，在 `render` 方法中你可以寫：

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

當 React 遇到一個訂製的 HTML 頁籤，該頁籤取用另一個組件（像這個範例中以 `< />` 包裹的組件名稱），它會在該頁籤的位置呈現該組件的標記。這應該說明 `App` 組件與 `Navbar`、`Dashboard` 和 `Footer` 之間的父子關係。

# --instructions--

在程式碼編輯器中，有一個簡單的函式組件叫做 `ChildComponent`，以及一個類別組件叫做 `ParentComponent`。將兩者組合在一起，方法是在 `ParentComponent` 中呈現 `ChildComponent`。請確保使用斜線關閉 `ChildComponent` 頁籤。

**注意：** `ChildComponent` 是用 ES6 箭頭函式定義的，因為這是在使用 React 時非常常見的做法。

# --hints--

React 組件應該傳回單一的 `div` 元素。

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

該組件應該傳回兩個巢狀元素。

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

該組件應將 `ChildComponent` 傳回為其第二個子組件。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ParentComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
