---
id: 5a24c314108439a4d4036163
title: 創建 React 組件
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

定義 React 組件的另一種方式是使用 ES6 的 `class` 語法。在以下範例中，`Kitten` 繼承自 `React.Component`：

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

這會創建一個 ES6 類別 `Kitten`，它繼承自 `React.Component` 類別。因此，`Kitten` 類別現在可以存取許多有用的 React 特性，例如區域的狀態和生命週期掛鉤。如果你還不熟悉這些字詞，不用擔心，後續的挑戰會更詳細涵蓋。還要注意 `Kitten` 類別內定義了一個 `constructor`，並呼叫了 `super()`。它使用 `super()` 來呼叫父類別的建構式，在這個例子中是 `React.Component`。建構式是在使用 `class` 關鍵字建立物件時用於初始化的特殊方法。呼叫組件的 `constructor` 時使用 `super` 並傳遞 `props` 是最佳實務，這確保組件能正確初始化。現在你只需知道這段程式碼包含這些是標準作法。很快你會看到建構式和 `props` 的其他用途。

# --instructions--

`MyComponent` 是在程式碼編輯器中使用類別語法定義的。完成撰寫 `render` 方法，使其傳回包含有文字 `Hello React!` 的 `h1` 的 `div` 元素。

# --hints--

React 組件應該傳回一個 `div` 元素。

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

傳回的 `div` 應該在其中呈現一個 `h1` 標題元素。

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

`h1` 標題元素應該包含字串 `Hello React!`。

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
