---
id: 5a24c314108439a4d4036184
title: 使用 If-Else 條件進行呈現
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

使用 JavaScript 控制你呈現的視圖的另一種應用是將呈現的元素與條件綁定。當條件為真時，呈現一個視圖。當條件為假時，呈現另一個視圖。你可以在 React 組件的 `render()` 方法中使用標準的 `if/else` 述語來做到這一點。

# --instructions--

MyComponent 在它的狀態中包含一個 `boolean`，用來追蹤你是否想在 UI 中顯示某個元素。`button` 會切換這個值的狀態。目前，它每次都呈現相同的 UI。請使用 `if/else` 述語重寫 `render()` 方法，使得如果 `display` 是 `true`，你就傳回目前的標記。否則，傳回不包含 `h1` 元素的標記。

**注意：**你必須撰寫 `if/else` 來通過測試。使用三元運算子在此不會通過。

# --hints--

`MyComponent` 應該存在並呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

當 `display` 設定為 `true` 時，應該呈現 `div`、`button` 和 `h1`。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

當 `display` 設定為 `false` 時，只有 `div` 和 `button` 應該呈現。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

呈現方法應該使用 `if/else` 述語來檢查 `this.state.display` 的條件。

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
