---
id: 5a24c314108439a4d403617e
title: 添加事件監聽器
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

`componentDidMount()` 方法也是為特定功能添加任何你需要的事件監聽器的最佳位置。React 提供了一個合成事件系統，該系統包裝了瀏覽器中存在的原生事件系統。這表示合成事件系統的行為在不同使用者的瀏覽器中完全相同——即使原生事件在不同瀏覽器之間可能有不同的行為。

你已經在使用一些這些合成事件處理常式，例如 `onClick()`。React 的合成事件系統非常適合用於你在 DOM 元素上管理的大多數互動。然而，如果你想將事件處理常式附加到文件或視窗物件，你必須直接這麼做。

# --instructions--

在 `componentDidMount()` 方法中附加一個 `keydown` 事件的監聽器，並讓這些事件觸發回呼 `handleKeyPress()`。你可以使用 `document.addEventListener()`，它將事件（以引號包住）作為第一個引數，回呼作為第二個引數。

然後，在 `componentWillUnmount()` 中，移除相同的事件監聽器。你可以將相同的引數傳遞給 `document.removeEventListener()`。在 React 組件被卸載和銷毀之前，使用這個生命週期方法來進行任何清理工作是良好的做法。移除事件監聽器就是其中一個清理動作的範例。

# --hints--

`MyComponent` 應該算繪一個包裹 `h1` 頁籤的 `div` 組件。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

應該在 `componentDidMount` 中將 `keydown` 監聽器附加到文件。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

應該在 `componentWillUnmount` 中從文件移除 `keydown` 監聽器。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

組件掛載後，按下 `enter` 應該會更新其狀態和算繪的 `h1` 頁籤。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
