---
id: 5a24c314108439a4d403616e
title: 使用 this.props 存取屬性
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

過去幾個挑戰涵蓋了將 props 傳遞給子組件的基本方法。但如果你傳遞 prop 的子組件是 ES6 類別組件，而不是無狀態函式組件，該怎麼辦？ES6 類別組件使用稍微不同的慣例來存取 props。

每當你在類別組件內部指涉該組件時，你會使用 `this` 關鍵字。若要存取類別組件內的 props，你需要在用來存取它的程式碼前加上 `this`。例如，若一個 ES6 類別組件有一個名為 `data` 的 prop，你會在 JSX 中寫 `{this.props.data}`。

# --instructions--

在父組件 `App` 中呈現 `Welcome` 組件的實體。在這裡，為 `Welcome` 指定一個名為 `name` 的 prop，並將其設值為一個字串。在子組件 `Welcome` 中，於 `strong` 標記內存取 `name` prop。

# --hints--

`App` 組件應該傳回單一的 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`App` 的子組件應該是 `Welcome` 組件。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

`Welcome` 組件應該有一個名為 `name` 的 prop。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

`Welcome` 組件應該在 `strong` 頁籤內顯示你作為 `name` prop 傳入的字串。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<App />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
