---
id: 5a24c314108439a4d403617c
title: 使用生命週期方法 `componentWillMount`
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

React 組件有幾個特殊的方法，提供在組件生命週期特定時刻執行動作的機會。這些稱為生命週期方法，或生命週期掛鉤，允許你在特定時間點捕捉組件。這可以是在它們被呈現之前、更新之前、接收 props 之前、卸載之前，等等。以下是一些主要生命週期方法的列表：`componentWillMount()`、`componentDidMount()`、`shouldComponentUpdate()`、`componentDidUpdate()`、`componentWillUnmount()`。接下來的幾課將涵蓋這些生命週期方法的一些基本使用案例。

**注意：** `componentWillMount` 生命週期方法將在未來的 16.X 版本中被棄用，並在 17 版本中移除。請在這篇 <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">文章</a> 中了解更多。

# --instructions--

當組件被掛載到 DOM 時，會在 `render()` 方法之前呼叫 `componentWillMount()` 方法。在 `componentWillMount()` 中將某些內容記錄到主控台－你可能想要打開瀏覽器主控台以查看輸出。

# --hints--

`MyComponent` 應該呈現一個 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

應該在 `componentWillMount` 中呼叫 `console.log`。

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
