---
id: 5a24c314108439a4d4036188
title: 根據 Props 有條件地呈現
challengeType: 6
forumTopicId: 301405
dashedName: render-conditionally-from-props
---

# --description--

到目前為止，你已經看到如何使用 `if/else`、`&&` 和三元運算子（`condition ? expressionIfTrue : expressionIfFalse`）來對要呈現的內容以及呈現的時間做出條件判斷。然而，還有一個重要的主題尚未討論，那就是讓你能將這些概念中的任意一個或全部與另一個強大的 React 特性結合起來：props。使用 props 來條件式呈現程式碼在 React 開發者中非常常見——也就是說，他們使用特定 prop 的值來自動決定要呈現什麼。

在這個挑戰中，你將設定一個子組件，根據 props 做出算繪決策。你也會使用三元運算子，但你可以看到在這個背景關係中，過去幾個挑戰所涵蓋的其他幾個概念同樣可能非常有用。

# --instructions--

程式碼編輯器有兩個為你部分定義的組件：一個名為 `GameOfChance` 的父組件，以及一個名為 `Results` 的子組件。它們用於創建一個簡單的遊戲，讓使用者按下按鈕以查看他們是否贏或輸。

首先，你需要一個簡單的表達式，每次執行時會隨機傳回不同的值。你可以使用 `Math.random()`。這個方法每次被呼叫時會傳回介於 `0`（包含）和 `1`（不包含）之間的值。因此，若要達到 50/50 的機率，在你的表達式中使用 `Math.random() >= .5`。從統計學角度來看，這個表達式有 50% 的機率傳回 `true`，另外 50% 傳回 `false`。在呈現方法中，將 `null` 取代為上述表達式，以完成變數宣告。

現在你有一個表達式，可以用來在程式碼中做出隨機決定。接下來你需要實作這個功能。將 `Results` 組件作為 `GameOfChance` 的子組件呈現，並將 `expression` 作為名為 `fiftyFifty` 的屬性傳入。在 `Results` 組件中，撰寫一個三元表達式，根據從 `GameOfChance` 傳入的 `fiftyFifty` 屬性，呈現帶有文字 `You Win!` 或 `You Lose!` 的 `h1` 元素。最後，確保 `handleClick()` 方法正確計算每一回合，讓使用者知道他們已經玩了多少次。這也能讓使用者知道組件確實已更新，以防他們連續贏或輸兩次。

# --hints--

`GameOfChance` 組件應該存在並呈現到頁面上。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('GameOfChance').length,
  1
);
```

`GameOfChance` 應該傳回單一的 `button` 元素。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('button').length,
  1
);
```

`GameOfChance` 應該傳回一個 `Results` 組件的單一實體，該組件有一個名為 `fiftyFifty` 的 prop。

```js
assert(
  Enzyme.mount(React.createElement(GameOfChance)).find('Results').length ===
    1 &&
    Enzyme.mount(React.createElement(GameOfChance))
      .find('Results')
      .props()
      .hasOwnProperty('fiftyFifty') === true
);
```

`GameOfChance` 狀態應該以 `counter` 屬性設定為值 `1` 來初始化。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).state().counter,
  1
);
```

當 `GameOfChance` 組件第一次被呈現到 DOM 時，應該傳回一個 `p` 元素，其內部文字為 `Turn: 1`。

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(GameOfChance)).find('p').text(),
  'Turn: 1'
);
```

每次按鈕被點擊時，計數器狀態應該增加 1，並且應該在 DOM 中呈現一個包含文字 `Turn: N` 的單一 `p` 元素，其中 `N` 是計數器狀態的值。

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    count: comp.state('counter'),
    text: comp.find('p').text()
  });
  const _1 = () => {
    simulate();
    return result();
  };
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _1_val = _1();
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  assert(
    _1_val.count === 2 &&
      _1_val.text === 'Turn: 2' &&
      _2_val.count === 3 &&
      _2_val.text === 'Turn: 3' &&
      _3_val.count === 4 &&
      _3_val.text === 'Turn: 4' &&
      _4_val.count === 5 &&
      _4_val.text === 'Turn: 5' &&
      _5_val.count === 6 &&
      _5_val.text === 'Turn: 6'
  );
})();
```

當 `GameOfChance` 組件第一次掛載到 DOM 時，以及之後每次按鈕被點擊時，都應該傳回一個單一的 `h1` 元素，該元素會隨機呈現 `You Win!` 或 `You Lose!`。注意：這可能會隨機失敗。如果發生這種情況，請再試一次。

```js
(() => {
  const comp = Enzyme.mount(React.createElement(GameOfChance));
  const simulate = () => {
    comp.find('button').simulate('click');
  };
  const result = () => ({
    h1: comp.find('h1').length,
    text: comp.find('h1').text()
  });
  const _1 = result();
  const _2 = () => {
    simulate();
    return result();
  };
  const _3 = () => {
    simulate();
    return result();
  };
  const _4 = () => {
    simulate();
    return result();
  };
  const _5 = () => {
    simulate();
    return result();
  };
  const _6 = () => {
    simulate();
    return result();
  };
  const _7 = () => {
    simulate();
    return result();
  };
  const _8 = () => {
    simulate();
    return result();
  };
  const _9 = () => {
    simulate();
    return result();
  };
  const _10 = () => {
    simulate();
    return result();
  };
  const _2_val = _2();
  const _3_val = _3();
  const _4_val = _4();
  const _5_val = _5();
  const _6_val = _6();
  const _7_val = _7();
  const _8_val = _8();
  const _9_val = _9();
  const _10_val = _10();
  const __text = new Set([
    _1.text,
    _2_val.text,
    _3_val.text,
    _4_val.text,
    _5_val.text,
    _6_val.text,
    _7_val.text,
    _8_val.text,
    _9_val.text,
    _10_val.text
  ]);
  const __h1 = new Set([
    _1.h1,
    _2_val.h1,
    _3_val.h1,
    _4_val.h1,
    _5_val.h1,
    _6_val.h1,
    _7_val.h1,
    _8_val.h1,
    _9_val.h1,
    _10_val.h1
  ]);
  assert(__text.size === 2 && __h1.size === 1);
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<GameOfChance />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    {/* Change code below this line */}
    return <h1></h1>;
    {/* Change code above this line */}
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      // Complete the return statement:
      return {
        counter: prevState
      }
    });
  }
  render() {
    const expression = null; // Change this line
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        {/* Change code below this line */}

        {/* Change code above this line */}
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```

# --solutions--

```jsx
// We want this to be deterministic for testing purposes.
const randomSequence = [true, false, false, true, true, false, false, true, true, false];
let index = 0;
const fiftyFifty = () => randomSequence[index++ % randomSequence.length];

class Results extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h1>{this.props.fiftyFifty ? 'You Win!' : 'You Lose!'}</h1>;
  }
}

class GameOfChance extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      counter: 1
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState(prevState => {
      return {
        counter: prevState.counter + 1
      }
    });
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Play Again</button>
        <Results fiftyFifty={fiftyFifty()} />
        <p>{'Turn: ' + this.state.counter}</p>
      </div>
    );
  }
}
```
