---
id: 5a24c314108439a4d4036182
title: 在 React 中添加行內樣式
challengeType: 6
forumTopicId: 301378
dashedName: add-inline-styles-in-react
---

# --description--

你可能已經注意到，在上一個挑戰中，除了將 `style` 屬性設為 JavaScript 物件之外，還有其他幾個與 HTML 行內樣式不同的語法差異。首先，某些 CSS 樣式屬性的名稱使用駝峰式大小寫。例如，上一個挑戰使用 `fontSize` 來設定字型大小，而不是 `font-size`。像 `font-size` 這樣帶有連字號的字是 JavaScript 物件屬性中無效的語法，因此 React 使用駝峰式大小寫。一般規則是，任何帶有連字號的樣式屬性在 JSX 中都以駝峰式大小寫書寫。

所有屬性值長度單位（例如 `height`、`width` 和 `fontSize`）預設為 `px`，除非另有指定。如果你想使用 `em`，例如，你需要將值和單位用引號括起來，如 `{fontSize: "4em"}`。除了預設為 `px` 的長度值外，所有其他屬性值都應該用引號括起來。

# --instructions--

如果你有大量的樣式，可以將樣式 `object` 指定給常數，以保持你的程式碼有條理。在檔案頂端將你的樣式常數宣告為全域變數。初始化 `styles` 常數並將一個包含三個樣式屬性及其值的 `object` 指定給它。給 `div` 指定顏色為 `purple`、字型大小為 `40`，以及邊框為 `2px solid purple`。然後將 `style` 屬性設值為 `styles` 常數。

# --hints--

`styles` 變數應該是一個具有三個屬性的 `物件`。

```js
assert(Object.keys(styles).length === 3);
```

`styles` 變數應該有一個 `color` 屬性設定為 `purple`。

```js
assert(styles.color === 'purple');
```

`styles` 變數應該有一個 `fontSize` 屬性設定為 `40` 的值。

```js
assert(styles.fontSize == 40);
```

`styles` 變數應該有一個 `border` 屬性設定為 `2px solid purple`。

```js
assert(styles.border === '2px solid purple');
```

該組件應該呈現一個 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return mockedComponent.type() === 'div';
  })()
);
```

`div` 元素的樣式應由 `styles` 物件定義。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.shallow(React.createElement(Colorful));
    return (
      mockedComponent.props().style.color === 'purple' &&
      mockedComponent.props().style.fontSize == 40 &&
      mockedComponent.props().style.border === '2px solid purple'
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Colorful />, document.getElementById('root'))
```

## --seed-contents--

```jsx
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={{color: "yellow", fontSize: 24}}>Style Me!</div>
    );
    // Change code above this line
  }
};
```

# --solutions--

```jsx
const styles = {
  color: "purple",
  fontSize: 40,
  border: "2px solid purple"
};
// Change code above this line
class Colorful extends React.Component {
  render() {
    // Change code below this line
    return (
      <div style={styles}>Style Me!</div>
    );
    // Change code above this line
  }
};
```
