---
id: 5a24c314108439a4d403617b
title: 將回呼作為 Props 傳遞
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

你可以將 `state` 作為 props 傳遞給子組件，但你不僅限於傳遞資料。你也可以將處理常式函式或任何在 React 組件上定義的方法傳遞給子組件。這就是你允許子組件與其父組件互動的方式。你像傳遞一般的 prop 一樣將方法傳遞給子組件。它會被指定一個名稱，而你可以在子組件的 `this.props` 下存取該方法名稱。

# --instructions--

程式碼編輯器中列出了三個組件。`MyApp` 組件是父組件，將呈現 `GetInput` 和 `RenderInput` 子組件。為你的 `MyApp` 中的呈現方法添加 `GetInput` 組件，然後將一個名為 `input` 的 prop 指定為 `MyApp` 的 `state` 中的 `inputValue`。同時創建一個名為 `handleChange` 的 prop，並將輸入處理常式 `handleChange` 傳遞給它。

接著，將 `RenderInput` 添加到 `MyApp` 的呈現方法中，然後創建一個名為 `input` 的 prop，並將 `state` 中的 `inputValue` 傳遞給它。完成後，你將能夠在 `GetInput` 組件中的 `input` 欄位輸入文字，該組件接著會透過 props 呼叫其父組件中的處理常式方法。這會更新父組件 `state` 中的輸入值，該值會作為 props 傳遞給兩個子組件。觀察資料如何在組件間流動，以及單一真實來源如何保持為父組件的 `state`。誠然，這個範例有點牽強，但應該能說明資料和回呼如何在 React 組件間傳遞。

# --hints--

`MyApp` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

`GetInput` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

`RenderInput` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

`GetInput` 組件應該接收 `MyApp` 狀態屬性 `inputValue` 作為 props，並包含一個修改 `MyApp` 狀態的 `input` 元素。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

`RenderInput` 組件應該接收 `MyApp` 狀態屬性 `inputValue` 作為 props。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyApp />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
