---
id: 5a24c314108439a4d403616b
title: 使用預設值 Props
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

React 也有為預設 props 設定選項。你可以將預設 props 指定為組件本身的屬性，React 會在必要時指派預設 prop。這讓你能夠指定當沒有明確提供值時，prop 的值應該是什麼。例如，如果你宣告 `MyComponent.defaultProps = { location: 'San Francisco' }`，你就定義了一個 location prop，預設值為字串 `San Francisco`，除非你另有指定。當 props 是 undefined 時，React 會指派預設 props，但如果你傳入 `null` 作為 prop 的值，它將保持為 `null`。

# --instructions--

程式碼編輯器顯示一個 `ShoppingCart` 組件。在此組件上定義預設 props，該預設 props 指定一個名為 `items` 且值為 `0` 的 prop。

# --hints--

`ShoppingCart` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

`ShoppingCart` 組件應該有一個預設的 prop 為 `{ items: 0 }`。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
