---
id: 5a24c314108439a4d403616a
title: 將陣列作為 Props 傳遞
challengeType: 6
forumTopicId: 301401
dashedName: pass-an-array-as-props
---

# --description--

上一個挑戰示範了如何將資訊從父組件傳遞到子組件作為 `props` 或屬性。這個挑戰探討如何將陣列作為 `props` 傳遞。要將陣列傳遞給 JSX 元素，必須將其視為 JavaScript 並用大括號包起來。

```jsx
<ParentComponent>
  <ChildComponent colors={["green", "blue", "red"]} />
</ParentComponent>
```

子組件隨後可以存取陣列屬性 `colors`。在存取該屬性時，可以使用像是 `join()` 的陣列方法。 

```jsx
const ChildComponent = (props) => <p>{props.colors.join(', ')}</p>
``` 

這將會將所有 `colors` 陣列項目以逗號分隔的字串連接起來，並產生：`<p>green, blue, red</p>`。稍後，我們將學習其他在 React 中呈現資料陣列的常用方法。

# --instructions--

程式碼編輯器中有 `List` 和 `ToDo` 組件。當從 `ToDo` 組件算繪每個 `List` 時，傳入一個 `tasks` 屬性，該屬性被指定為待辦任務的陣列，例如 `["walk dog", "workout"]`。然後在 `List` 組件中存取這個 `tasks` 陣列，並在 `p` 元素中顯示其值。使用 `join(", ")` 將 `props.tasks` 陣列以逗號分隔的列表形式顯示在 `p` 元素中。今天的列表應該至少有 2 個任務，明天的應該至少有 3 個任務。

# --hints--

`ToDo` 組件應該傳回一個單一的外層 `div`。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().type() === 'div';
  })()
);
```

`ToDo` 組件的第三個子項目應該是 `List` 組件的實體。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(2).name() === 'List';
  })()
);
```

`ToDo` 組件的第五個子項目應該是 `List` 組件的實體。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.children().first().childAt(4).name() === 'List';
  })()
);
```

兩個 `List` 組件的實體都應該有一個名為 `tasks` 的屬性，且 `tasks` 應該是陣列型別。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      Array.isArray(mockedComponent.find('List').get(0).props.tasks) &&
      Array.isArray(mockedComponent.find('List').get(1).props.tasks)
    );
  })()
);
```

第一個表述今日任務的 `List` 組件應該有 2 個或以上的條款。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(0).props.tasks.length >= 2;
  })()
);
```

第二個表述明天任務的 `List` 組件應該有 3 個或更多項目。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return mockedComponent.find('List').get(1).props.tasks.length >= 3;
  })()
);
```

`List` 組件應該在 `p` 頁籤中呈現來自 `tasks` 屬性的值。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ToDo));
    return (
      mockedComponent
        .find('p')
        .get(0)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(0)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',') &&
      mockedComponent
        .find('p')
        .get(1)
        .props.children.replace(/\s*,\s*/g, ',') ===
        mockedComponent
          .find('List')
          .get(1)
          .props.tasks.join(',')
          .replace(/\s*,\s*/g, ',')
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ToDo />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const List = (props) => {
  { /* Change code below this line */ }
  return <p>{}</p>
  { /* Change code above this line */ }
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        { /* Change code below this line */ }
        <List/>
        <h2>Tomorrow</h2>
        <List/>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const List= (props) => {
  return <p>{props.tasks.join(', ')}</p>
};

class ToDo extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>To Do Lists</h1>
        <h2>Today</h2>
        <List tasks={['study', 'exercise']} />
        <h2>Tomorrow</h2>
        <List tasks={['call Sam', 'grocery shopping', 'order tickets']} />
      </div>
    );
  }
};
```
