---
id: 5a24c314108439a4d403616c
title: 改寫預設值 Props
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

設定預設 props 的能力是 React 中一個有用的特性。改寫預設 props 的方法是為組件明確地設定 prop 值。

# --instructions--

`ShoppingCart` 組件現在會呈現一個子組件 `Items`。這個 `Items` 組件有一個預設的屬性 `quantity`，其值為整數 `0`。將 `quantity` 的值傳入 `10`，以改寫預設屬性。

**注意：** 請記住，為組件添加 prop 的語法看起來類似於你添加 HTML 屬性的方式。然而，由於 `quantity` 的值是整數，它不會放在引號中，而應該包裹在大括號內。例如，`{100}`。這個語法告訴 JSX 直接將大括號內的值解釋為 JavaScript。

# --hints--

組件 `ShoppingCart` 應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

組件 `Items` 應該呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

`Items` 組件應該從 `ShoppingCart` 組件接收一個屬性為 `{ quantity: 10 }`。

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
