---
id: 5a24c314108439a4d4036170
title: 創建有狀態的組件
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

React 中最重要的主題之一是 `state`。`state` 包含你的應用程式需要知道且會隨時間改變的任何資料。你希望你的應用程式能對 `state` 變化做出回應，並在必要時呈現更新的使用者介面。React 為現代網頁應用程式的 `state` 管理提供了一個很好的解決方案。

你在 React 組件中透過在其 `constructor` 的組件類別上宣告 `state` 屬性來創建狀態。這會在組件被建立時初始化 `state`。`state` 屬性必須設定為 JavaScript `object`。宣告方式如下：

```jsx
this.state = {

}
```

你可以在組件的整個生命週期中存取 `state` 物件。你可以更新它、在你的 UI 中呈現它，並將它作為 props 傳遞給子組件。`state` 物件可以根據你的需求變得非常複雜或非常簡單。請注意，你必須透過繼承 `React.Component` 來創建類別組件，才能像這樣創建 `state`。

# --instructions--

程式碼編輯器中有一個組件試圖從它的 `state` 呈現 `firstName` 屬性。然而，沒有定義 `state`。請在 `constructor` 中初始化組件的 `state`，並將你的名字指定到 `firstName` 屬性。

# --hints--

`StatefulComponent` 應該存在並呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent` 應該呈現一個 `div` 和一個 `h1` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

`StatefulComponent` 的狀態應該以一個屬性 `firstName` 初始化，該屬性設定為一個字串。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

`StatefulComponent` 狀態中的屬性 `firstName` 應該在 `h1` 元素中呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<StatefulComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
