---
id: 5a24c314108439a4d403616f
title: 檢視在無狀態函式（程式）組件中使用 Props
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

除了最後一個挑戰，你一直在將 props 傳遞給無狀態函式（程式）組件。這些組件的行為就像純函式（程式）。它們接受 props 作為輸入，並且每次傳入相同的 props 時，都會傳回相同的視圖。你可能會想知道什麼是 state，下一個挑戰將會更詳細地涵蓋它。在此之前，這裡是組件術語的複習。

*無狀態函式組件* 是你撰寫的任何接受 props 並傳回 JSX 的函式（程式）。另一方面，*無狀態組件* 是繼承自 `React.Component` 的類別，但不使用內部狀態（涵蓋於下一個挑戰）。最後，*有狀態組件* 是會維護自身內部狀態的類別組件。你可能會看到有狀態組件簡稱為組件或 React 組件。

一個常見的樣式是盡量減少有狀態性，並在可能的情況下創建無狀態的函式（程式）組件。這有助於將你的狀態管理限制在應用程式的特定區域。反過來，這透過讓你更容易追蹤狀態變更如何影響其行為，提升你的應用程式的開發和維護。

# --instructions--

程式碼編輯器有一個 `CampSite` 組件，它會將 `Camper` 組件作為子組件來呈現。定義 `Camper` 組件並為它指定預設屬性為 `{ name: 'CamperBot' }`。在 `Camper` 組件內，算繪你想要的任何程式碼，但請確保有一個 `p` 元素，該元素只包含作為 `prop` 傳入的 `name` 值。最後，在 `Camper` 組件上定義 `propTypes`，要求必須提供 `name` 作為屬性，並驗證它的型別為 `string`。

# --hints--

`CampSite` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

`Camper` 組件應該要呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

`Camper` 組件應該包含預設 props，將字串 `CamperBot` 指定給鍵 `name`。

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

`Camper` 組件應該包含 prop 型別，要求 `name` prop 必須是 `string` 型別。

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

`Camper` 組件應該包含一個 `p` 元素，該元素僅包含來自 `name` prop 的文字。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
var PropTypes = {
   string: { isRequired: true }
};
```

## --after-user-code--

```jsx
ReactDOM.render(<CampSite />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
