---
id: 5a24c314108439a4d403618b
title: 為兄弟元素添加唯一的 Key 屬性
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

上一個挑戰展示了如何使用 `map` 方法根據使用者輸入動態呈現多個元素。然而，那個範例中缺少了一個重要的部分。當你創建一個元素陣列時，每個元素都需要設定一個唯一值的 `key` 屬性。React 使用這些鍵來追蹤哪些條款被新增、變更或移除。當列表以任何方式被修改時，這有助於使重新呈現過程更高效。  

**注意：** 鍵只需要在同層元素之間唯一，無需在你的應用程式中全域的唯一。

# --instructions--

程式碼編輯器中有一個包含一些前端框架的陣列，以及一個名為 `Frameworks()` 的無狀態函式組件。`Frameworks()` 需要將陣列映射成一個無序列表，就像上一個挑戰一樣。完成撰寫 `map` 回呼，為 `frontEndFrameworks` 陣列中的每個框架傳回一個 `li` 元素。這次，請確保為每個 `li` 元素設定一個唯一值的 `key` 屬性。`li` 元素也應該包含來自 `frontEndFrameworks` 的文字。

通常，你會想要將 key 設為能唯一識別被算繪元素的值。作為最後手段，可以使用陣列索引，但通常你應該嘗試使用唯一識別。

# --hints--

`Frameworks` 組件應該存在並呈現到頁面上。

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` 應該呈現一個 `h1` 元素。

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` 應該呈現一個 `ul` 元素。

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

`ul` 頁籤應該呈現 6 個子 `li` 元素。

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

每個列表項目元素都應該有一個唯一的 `key` 屬性。

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

每個列表項目元素應包含來自 `frontEndFrameworks` 的文字。

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Frameworks />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Frontend JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Frontend JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
