---
id: 5a24c314108439a4d4036160
title: 在 JSX 中定義 HTML 類別
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

現在你已經開始熟悉撰寫 JSX，可能會想知道它與 HTML 有何不同。

到目前為止，HTML 和 JSX 可能看起來完全相同。

JSX 的一個主要差異是你不能再使用字 `class` 來定義 HTML 類別。這是因為 `class` 是 JavaScript 的保留字。取而代之，JSX 使用 `className`。

事實上，JSX 中所有 HTML 屬性和事件引用的命名慣例都改為駝峰式。例如，JSX 中的點擊事件是 `onClick`，而不是 `onclick`。同樣地，`onchange` 變成 `onChange`。雖然這是個細微的差別，但在未來必須牢記這一點。

# --instructions--

將 `myDiv` 類別套用到 JSX 程式碼中提供的 `div`。

# --hints--

常數 `JSX` 應該傳回一個 `div` 元素。

```js
assert.strictEqual(JSX.type, 'div');
```

`div` 應該有一個 `myDiv` 的類別。

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(JSX, document.getElementById('root'))
```

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
