---
id: 5a24c314108439a4d4036172
title: 在使用者介面中呈現狀態的另一種方式
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

在組件中還有另一種存取 `state` 的方式。在 `render()` 方法中，於 `return` 述語之前，你可以直接撰寫 JavaScript。例如，你可以宣告函式、存取來自 `state` 或 `props` 的資料、對這些資料進行運算，等等。然後，你可以將任何資料指定給變數，而你可以在 `return` 述語中存取這些變數。

# --instructions--

在 `MyComponent` 的呈現方法中，定義一個名為 `name` 的 `const`，並將它設定為組件 `state` 中的 name 值。因為你可以直接在程式碼的這部分撰寫 JavaScript，所以你不必將此引用包在大括號中。

接著，在 `return` 述語中，使用變數 `name` 在 `h1` 頁籤中呈現此值。請記得，你需要在 `return` 述語中使用 JSX 語法（大括號用於 JavaScript）。

# --hints--

`MyComponent` 應該在其 state 中有一個 key 為 `name`、值為 `freeCodeCamp`。

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` 應該在單一 `div` 中呈現一個 `h1` 標題元素。

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

算繪後的 `h1` 頁籤應該包含對 `{name}` 的引用。

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

算繪後的 `h1` 標題組件元素應包含從組件狀態算繪的文字。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
