---
id: 5a24c314108439a4d4036165
title: 使用 React 呈現巢狀組件
challengeType: 6
forumTopicId: 301420
dashedName: use-react-to-render-nested-components
---

# --description--

上一個挑戰展示了一種簡單的方式來組合兩個組件，但你可以用 React 組合組件的方式有很多種。

組件組合是 React 強大的特性之一。當你使用 React 時，重要的是開始以組件的角度來思考你的使用者介面，就像上一個挑戰中的 App 範例一樣。你將你的使用者介面拆解成基本的建造區塊，而這些部分就成為組件。這有助於將負責使用者介面的程式碼與負責處理應用程式邏輯的程式碼分離。它能大幅簡化複雜專案的開發與維護。

# --instructions--

在程式碼編輯器中定義了兩個函式組件，分別稱為 `TypesOfFruit` 和 `Fruits`。將 `TypesOfFruit` 組件取出並組合，或*巢狀*，在 `Fruits` 組件內。然後將 `Fruits` 組件巢狀在 `TypesOfFood` 組件內。結果應該是一個子組件，巢狀在父組件內，而該父組件又巢狀在另一個父組件內！

# --hints--

`TypesOfFood` 組件應該傳回單一的 `div` 元素。

```js
assert(Enzyme.shallow(React.createElement(TypesOfFood)).type() === 'div');
```

`TypesOfFood` 組件應該傳回 `Fruits` 組件。

```js
assert(
  Enzyme.shallow(React.createElement(TypesOfFood)).props().children[1].type
    .name === 'Fruits'
);
```

`Fruits` 組件應該傳回 `TypesOfFruit` 組件。

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('h2').html() ===
    '<h2>Fruits:</h2>'
);
```

`TypesOfFruit` 組件應該傳回 `h2` 和 `ul` 元素。

```js
assert(
  Enzyme.mount(React.createElement(TypesOfFood)).find('ul').text() ===
    'ApplesBlueberriesStrawberriesBananas'
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<TypesOfFood />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }

      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const TypesOfFruit = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <ul>
        <li>Apples</li>
        <li>Blueberries</li>
        <li>Strawberries</li>
        <li>Bananas</li>
      </ul>
    </div>
  );
};

const Fruits = () => {
  return (
    <div>
      { /* Change code below this line */ }
        <TypesOfFruit />
      { /* Change code above this line */ }
    </div>
  );
};

class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        { /* Change code below this line */ }
        <Fruits />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
