---
id: 5a24c314108439a4d403617d
title: 使用生命週期方法 `componentDidMount`
challengeType: 6
forumTopicId: 301422
dashedName: use-the-lifecycle-method-componentdidmount
---

# --description--

大多數網頁開發者在某個階段都需要呼叫 API 端點以擷取資料。如果你正在使用 React，了解在哪裡執行此操作非常重要。

React 的最佳實務是在生命週期方法 `componentDidMount()` 中放置 API 呼叫或任何對你的伺服端的呼叫。此方法會在組件掛載到 DOM 後被呼叫。此處對 `setState()` 的任何呼叫都會觸發組件的重新呈現。當你在此方法中呼叫 API，並用 API 傳回的資料設定你的狀態時，一旦收到資料，它會自動觸發更新。

# --instructions--

在 `componentDidMount()` 中有一個模擬的 API 呼叫。它在 2.5 秒後設定狀態，以模擬呼叫伺服器來取得資料。此範例請求網站目前的總作用中使用者數。在呈現方法中，將 `activeUsers` 的值在 `h1` 中的文字 `Active Users:` 後呈現。觀察預覽中的變化，並隨意更改逾時時間以查看不同效果。

# --hints--

`MyComponent` 應該算繪一個包裹 `h1` 頁籤的 `div` 組件。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

組件狀態應該在 `componentDidMount` 中使用 timeout 函式更新。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return new RegExp('setTimeout(.|\n)+setState(.|\n)+activeUsers').test(
      String(mockedComponent.instance().componentDidMount)
    );
  })()
);
```

`h1` 頁籤應該呈現 `MyComponent` 狀態中的 `activeUsers` 值。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ activeUsers: 1237 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ activeUsers: 1000 });
    return mockedComponent.find('h1').text();
  };
  assert(new RegExp('1237').test(first()) && new RegExp('1000').test(second()));
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        {/* Change code below this line */}
        <h1>Active Users: </h1>
        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      activeUsers: null
    };
  }
  componentDidMount() {
    setTimeout(() => {
      this.setState({
        activeUsers: 1273
      });
    }, 2500);
  }
  render() {
    return (
      <div>
        <h1>Active Users: {this.state.activeUsers}</h1>
      </div>
    );
  }
}
```
