---
id: 5a24c314108439a4d4036167
title: 將類別組件呈現到 DOM
challengeType: 6
forumTopicId: 301404
dashedName: render-a-class-component-to-the-dom
---

# --description--

你可能還記得在之前的挑戰中使用 ReactDOM API 將 JSX 元素呈現到 DOM。呈現 React 組件的處理程序看起來會非常相似。過去幾個挑戰專注於組件和組合，因此呈現是由系統在背後為你完成的。然而，除非你呼叫 ReactDOM API，否則你撰寫的任何 React 程式碼都不會呈現到 DOM。

這裡是語法的複習：`ReactDOM.render(componentToRender, targetNode)`。第一個引數是你想要呈現的 React 組件。第二個引數是你想要在其中呈現該組件的 DOM 節點。

React 組件傳入 `ReactDOM.render()` 的方式與 JSX 元素略有不同。對於 JSX 元素，你傳入你想要呈現的元素名稱。然而，對於 React 組件，你需要使用與呈現巢狀組件相同的語法，例如 `ReactDOM.render(<ComponentToRender />, targetNode)`。你對 ES6 類別組件和函式組件都使用此語法。

# --instructions--

`Fruits` 和 `Vegetables` 組件都已在幕後為你定義。將這兩個組件作為 `TypesOfFood` 組件的子組件來呈現，然後將 `TypesOfFood` 呈現到 DOM。你可以使用一個帶有 `id='challenge-node'` 的 `div`。

# --hints--

`TypesOfFood` 組件應該傳回單一的 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`TypesOfFood` 組件應該在 `h1` 元素之後呈現 `Fruits` 組件。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(1).name() === 'Fruits';
  })()
);
```

`TypesOfFood` 組件應該在 `Fruits` 之後呈現 `Vegetables` 組件。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(TypesOfFood));
    return mockedComponent.children().childAt(2).name() === 'Vegetables';
  })()
);
```

`TypesOfFood` 組件應該在具有 id 為 `challenge-node` 的 `div` 中呈現到 DOM。

```js
assert(
  (function () {
    const html = document.getElementById('challenge-node').childNodes[0]
      .innerHTML;
    return (
      html.includes(
        '<div><h2>Fruits:</h2><h4>Non-Citrus:</h4><ul><li>Apples</li><li>Blueberries</li><li>Strawberries</li><li>Bananas</li></ul><h4>Citrus:</h4><ul><li>Lemon</li><li>Lime</li><li>Orange</li><li>Grapefruit</li></ul></div>'
      ) &&
      html.includes(
        '<div><h2>Vegetables:</h2><ul><li>Brussel Sprouts</li><li>Broccoli</li><li>Squash</li></ul></div>'
      )
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
const Fruits = () => {
  return (
    <div>
      <h2>Fruits:</h2>
      <h4>Non-Citrus:</h4>
        <ul>
          <li>Apples</li>
          <li>Blueberries</li>
          <li>Strawberries</li>
          <li>Bananas</li>
        </ul>
      <h4>Citrus:</h4>
        <ul>
          <li>Lemon</li>
          <li>Lime</li>
          <li>Orange</li>
          <li>Grapefruit</li>
        </ul>
    </div>
  );
};
const Vegetables = () => {
  return (
    <div>
      <h2>Vegetables:</h2>
      <ul>
        <li>Brussel Sprouts</li>
        <li>Broccoli</li>
        <li>Squash</li>
      </ul>
    </div>
  );
};
```

## --seed-contents--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
```

# --solutions--

```jsx
class TypesOfFood extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>Types of Food:</h1>
        {/* Change code below this line */}
          <Fruits />
           <Vegetables />
         {/* Change code above this line */}
      </div>
    );
  }
};

// Change code below this line
ReactDOM.render(<TypesOfFood />, document.getElementById('challenge-node'));
```
