---
id: 5a24c314108439a4d4036161
title: 了解自我關閉的 JSX 頁籤
challengeType: 6
forumTopicId: 301396
dashedName: learn-about-self-closing-jsx-tags
---

# --description--

到目前為止，你已經看到 JSX 在使用 `className` 與 `class` 定義 HTML 類別時的主要差異。

JSX 與 HTML 不同的另一個重要方式是在自我關閉頁籤的概念上。

在 HTML 中，幾乎所有的頁籤都有開啟和關閉頁籤：`<div></div>`；關閉頁籤總是在你要關閉的頁籤名稱前面加上斜線。然而，HTML 中有特殊的實體稱為 <dfn>void elements</dfn>，或是不需要在另一個元素開始前同時有開啟和關閉頁籤的元素。

例如換行標記可以寫成 `<br>` 或 `<br />`，但絕不應該寫成 `<br></br>`，因為它不包含任何內容。

在 JSX 中，規則有些不同。任何 JSX 元素都可以使用自我關閉的標記來撰寫，且每個元素都必須被關閉。換行標記，<dfn>例如</dfn>，必須始終寫成 `<br />`，才能成為有效且可轉譯的 JSX。另一方面，`<div>` 可以寫成 `<div />` 或 `<div></div>`。差別在於第一種語法版本中無法在 `<div />` 中包含任何內容。你會在後續的挑戰中看到，當算繪 React 組件時，這種語法非常有用。

# --instructions--

修正程式碼編輯器中的錯誤，使其成為有效的 JSX 並成功轉譯。請確保你不更改任何內容－你只需要在需要的地方關閉標籤。

# --hints--

常數 `JSX` 應該傳回一個 `div` 元素。

```js
assert.strictEqual(JSX.type, 'div');
```

`div` 應該包含一個 `br` 頁籤。

```js
assert(Enzyme.shallow(JSX).find('br').length === 1);
```

`div` 應該包含一個 `hr` 頁籤。

```js
assert(Enzyme.shallow(JSX).find('hr').length === 1);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(JSX, document.getElementById('root'))
```

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h2>Welcome to React!</h2> <br >
    <p>Be sure to close all tags!</p>
    <hr >
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div>
  <h2>Welcome to React!</h2> <br />
  <p>Be sure to close all tags!</p>
  <hr />
</div>
);
```
