---
id: 5a24c314108439a4d4036169
title: 將 Props 傳遞給無狀態函式組件
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

先前的挑戰涵蓋了許多關於建立和組合 JSX 元素、函式組件以及 ES6 風格類別組件在 React 中的內容。有了這個基礎，現在是時候來看看 React 中另一個非常常見的特性：**props**。在 React 中，你可以將 props，或稱屬性，傳遞給子組件。假設你有一個 `App` 組件，它呈現一個名為 `Welcome` 的子組件，而 `Welcome` 是一個無狀態函式組件。你可以透過以下寫法將 `user` 屬性傳遞給 `Welcome`：

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

你使用由你創建且被 React 支援的**訂製 HTML 屬性**來傳遞給組件。在此情況下，創建的屬性 `user` 被傳遞給組件 `Welcome`。由於 `Welcome` 是無狀態函式組件，它可以像這樣存取此值：

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

標準做法是將此值稱為 `props`，當處理無狀態函式組件時，你基本上會將它視為一個傳回 JSX 的函式引數。你可以在函式主體中存取該引數的值。使用類別組件時，你會發現這有些不同。

# --instructions--

程式碼編輯器中有 `Calendar` 和 `CurrentDate` 組件。當從 `Calendar` 組件算繪 `CurrentDate` 時，傳入一個 `date` 屬性，其值指定為 JavaScript 的 `Date` 物件所取得的當前日期。然後在 `CurrentDate` 組件中存取這個 `prop`，並將其值顯示在 `p` 標記內。請注意，為了讓 `prop` 值能以 JavaScript 求值，必須用大括號包住，例如 `date={Date()}`。

# --hints--

`Calendar` 組件應該傳回單一的 `div` 元素。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

`Calendar` 組件的第二個子項目應該是 `CurrentDate` 組件。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

`CurrentDate` 組件應該有一個名為 `date` 的 prop。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

`CurrentDate` 的 `date` 屬性應包含一個字串。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

`date` 屬性應該透過呼叫 `Date()` 來產生。

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

`CurrentDate` 組件應該在 `p` 頁籤中呈現來自 `date` 屬性的值。

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Calendar />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
