---
id: 5a24c314108439a4d4036185
title: 使用 && 來撰寫更簡潔的條件判斷
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

`if/else` 述語在上一個挑戰中有效，但有一種更簡潔的方法可以達成相同的結果。想像你在一個組件中追蹤多個條件，並且你想根據每個條件呈現不同的元素。如果你寫很多 `else if` 述語來傳回略有不同的使用者介面，你可能會重複程式碼，這會留下錯誤的空間。相反地，你可以使用 `&&` 邏輯運算子以更簡潔的方式執行條件邏輯。這是可能的，因為你想檢查條件是否為 `true`，如果是，就傳回一些標記。以下是一個範例：

```jsx
{condition && <p>markup</p>}
```

如果 `condition` 是 `true`，則會傳回標記。 如果 `condition` 是 `false`，則在評估 `condition` 後，操作會立即傳回 `false` 並且不傳回任何內容。你可以直接在你的 JSX 中包含這些述語，並且透過在每個述語後面寫 `&&` 來串接多個條件。這讓你能在 `render()` 方法中處理更複雜的條件邏輯，而不必重複大量程式碼。

# --instructions--

再次解決前面的範例，使得 `h1` 只有在 `display` 為 `true` 時才呈現，但使用 `&&` 邏輯運算子而非 `if/else` 述語。

# --hints--

`MyComponent` 應該存在並呈現。

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

當 `display` 設定為 `true` 時，應該呈現 `div`、`button` 和 `h1`。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

當 `display` 設定為 `false` 時，只有 `div` 和 `button` 應該呈現。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

呈現方法應該使用 `&&` 邏輯運算子來檢查 `this.state.display` 的條件。

```js
assert(code.includes('&&'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
