---
id: 5a24c314108439a4d4036171
title: 使用者介面中的呈現狀態
challengeType: 6
forumTopicId: 301409
dashedName: render-state-in-the-user-interface
---

# --description--

一旦你定義了組件的初始狀態，你就可以在被算繪的使用者介面中呈現它的任何部分。如果組件是有狀態的，它將始終能在其 `render()` 方法中存取 `state` 中的資料。你可以使用 `this.state` 來存取資料。

如果你想在 `return` 的呈現方法中存取狀態值，你必須將該值包在大括號中。

`state` 是 React 組件中最強大的特性之一。它允許你追蹤應用程式中重要的資料，並根據這些資料的變化來呈現使用者介面。如果你的資料改變，使用者介面也會改變。React 使用所謂的虛擬 DOM，在幕後追蹤變化。當 `state` 資料更新時，它會觸發使用該資料的組件重新呈現——包括接收該資料作為 prop 的子組件。React 會更新實際的 DOM，但只在必要的地方更新。這表示你不必擔心改變 DOM，你只需宣告使用者介面應該呈現的樣子。

請注意，如果你使組件具有狀態，其他組件不會知道它的 `state`。它的 `state` 是完全封裝的，或是區域的，除非你將狀態資料作為 `props` 傳遞給子組件。這種封裝的 `state` 概念非常重要，因為它允許你撰寫特定邏輯，然後將該邏輯封裝並隔離在程式碼中的一個位置。

# --instructions--

在程式碼編輯器中，`MyComponent` 已經是有狀態的。請在組件的算繪方法中定義一個 `h1` 頁籤，該頁籤呈現組件狀態中 `name` 的值。

**注意：** `h1` 應該只呈現來自 `state` 的值，且不包含其他內容。在 JSX 中，任何用大括號 `{ }` 包住的程式碼都會被視為 JavaScript。所以若要存取來自 `state` 的值，只需將引用用大括號包起來。

# --hints--

`MyComponent` 應該在其 state 中有一個 key 為 `name`、值為 `freeCodeCamp`。

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` 應該在單一 `div` 中呈現一個 `h1` 標題元素。

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

算繪的 `h1` 標題元素應該只包含從組件狀態算繪的文字。

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  const getValue = firstValue.replace(/\s/g, '');
  assert(getValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{this.state.name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
