---
id: 5a24c314108439a4d4036179
title: 創建受控表單
challengeType: 6
forumTopicId: 301384
dashedName: create-a-controlled-form
---

# --description--

最後一個挑戰顯示 React 可以控制某些元素如 `input` 和 `textarea` 的內部狀態，這使它們成為受控組件。這也適用於其他表單元素，包括一般的 HTML `form` 元素。

# --instructions--

`MyForm` 組件設置了一個空的 `form`，並帶有一個提交處理常式。當表單提交時，將會呼叫該提交處理常式。

我們已經添加了一個按鈕用來提交表單。你可以看到它的 `type` 設定為 `submit`，表示它是控制該表單的按鈕。將 `input` 元素添加到 `form` 中，並像上一個挑戰一樣設定它的 `value` 和 `onChange()` 屬性。接著你應該完成 `handleSubmit` 方法，使其將組件狀態屬性 `submit` 設定為區域 `state` 中目前的輸入值。

**注意：**你也必須在送出處理常式中呼叫 `event.preventDefault()`，以防止會重新整理網頁的預設表單送出行為。為了露營者的方便，這裡已停用預設行為，以防止重新整理導致挑戰程式碼被重設。

最後，在 `form` 之後創建一個 `h1` 頁籤，該頁籤呈現組件的 `state` 中的 `submit` 值。然後你可以在表單中輸入並點擊按鈕（或按下 enter 鍵），你應該會看到你的輸入被呈現在頁面上。

# --hints--

`MyForm` 應該傳回一個包含 `form` 和 `h1` 頁籤的 `div` 元素。該表單應該包含一個 `input` 和一個 `button`。

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyForm));
    return (
      mockedComponent.find('div').children().find('form').length === 1 &&
      mockedComponent.find('div').children().find('h1').length === 1 &&
      mockedComponent.find('form').children().find('input').length === 1 &&
      mockedComponent.find('form').children().find('button').length === 1
    );
  })()
);
```

`MyForm` 的狀態應該以 `input` 和 `submit` 屬性初始化，兩者皆設定為空字串。

```js
assert(
  Enzyme.mount(React.createElement(MyForm)).state('input') === '' &&
    Enzyme.mount(React.createElement(MyForm)).state('submit') === ''
);
```

在 `input` 元素中輸入應該會更新組件狀態的 `input` 屬性。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    return mockedComponent.state('input');
  };
  const _2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return {
      state: mockedComponent.state('input'),
      inputVal: mockedComponent.find('input').props().value
    };
  };
  const before = _1();
  const after = _2();
  assert(
    before === '' &&
      after.state === 'TestInput' &&
      after.inputVal === 'TestInput'
  );
})();
```

提交表單時應該執行 `handleSubmit`，該函式應該將狀態中的 `submit` 屬性設定為目前輸入的值。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'SubmitInput' } });
    return mockedComponent.state('submit');
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.state('submit');
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'SubmitInput');
})();
```

`handleSubmit` 應該呼叫 `event.preventDefault`

```js
assert(
  __helpers.isCalledWithNoArgs(
    'event.preventDefault',
    MyForm.prototype.handleSubmit.toString()
  )
);
```

`h1` 標題元素應該呈現組件狀態中 `submit` 欄的值。

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyForm));
  const _1 = () => {
    mockedComponent.setState({ input: '' });
    mockedComponent.setState({ submit: '' });
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return mockedComponent.find('h1').text();
  };
  const _2 = () => {
    mockedComponent.find('form').simulate('submit');
    return mockedComponent.find('h1').text();
  };
  const before = _1();
  const after = _2();
  assert(before === '' && after === 'TestInput');
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyForm />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          {/* Change code below this line */}

          {/* Change code above this line */}
          <button type='submit'>Submit!</button>
        </form>
        {/* Change code below this line */}

        {/* Change code above this line */}
      </div>
    );
  }
}
```

# --solutions--

```jsx
class MyForm extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      submit: ''
    };
    this.handleChange = this.handleChange.bind(this);
    this.handleSubmit = this.handleSubmit.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  handleSubmit(event) {
    event.preventDefault();
    this.setState(state => ({
      submit: state.input
    }));
  }
  render() {
    return (
      <div>
        <form onSubmit={this.handleSubmit}>
          <input value={this.state.input} onChange={this.handleChange} />
          <button type='submit'>Submit!</button>
        </form>
        <h1>{this.state.submit}</h1>
      </div>
    );
  }
}
```
