---
id: 5a24c314108439a4d4036173
title: Set State with this.setState
challengeType: 6
forumTopicId: 301412
dashedName: set-state-with-this-setstate
---

# --description--

Le sfide precedenti hanno trattato il componente `state` e come inizializzare lo stato nel `constructor`. Esiste anche un modo per modificare lo `state` del componente. React fornisce un metodo per aggiornare lo stato del componente chiamato `setState`. Chiami il metodo `setState` all’interno della tua classe componente in questo modo: `this.setState()`, passando un oggetto con coppie chiave-valore. Le chiavi sono le proprietà dello `state` e i valori sono i dati aggiornati dello `state`. Per esempio, se stessimo memorizzando una `username` nello `state` e volessimo aggiornarla, sarebbe così:

```jsx
this.setState({
  username: 'Lewis'
});
```

React si aspetta che tu non modifichi mai direttamente lo `state`, ma che usi sempre `this.setState()` quando avvengono cambiamenti di stato. Inoltre, devi sapere che React può raggruppare più aggiornamenti di stato per migliorare le prestazioni. Questo significa che gli aggiornamenti di stato tramite il metodo `setState` possono essere asincroni. Esiste una sintassi alternativa per il metodo `setState` che offre una soluzione a questo problema. È raramente necessaria, ma è bene tenerla a mente! Consulta il nostro <a href="https://www.freecodecamp.org/news/what-is-state-in-react-explained-with-examples/" target="_blank" rel="noopener noreferrer nofollow">articolo su React</a> per ulteriori dettagli.

# --instructions--

C’è un elemento `button` nell’editor di codice che ha un gestore `onClick()`. Questo gestore viene attivato quando il `button` riceve un evento di clic nel browser, e esegue il metodo `handleClick` definito su `MyComponent`. All’interno del metodo `handleClick`, aggiorna il componente `state` usando `this.setState()`. Imposta la proprietà `name` in `state` uguale alla stringa `React Rocks!`.

Clicca il pulsante e osserva l’aggiornamento dello stato renderizzato. Non preoccuparti se non capisci completamente come funziona il codice del gestore clic in questo momento. Verrà trattato nelle sfide successive.

# --hints--

Lo stato di `MyComponent` dovrebbe inizializzarsi con la coppia chiave-valore `{ name: Initial State }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'Initial State'
);
```

`MyComponent` dovrebbe renderizzare un elemento intestazione `h1`.

```js
assert(Enzyme.mount(React.createElement(MyComponent)).find('h1').length === 1);
```

L’elemento intestazione `h1` renderizzato dovrebbe contenere testo renderizzato dallo stato del componente.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(/<h1>TestName<\/h1>/.test(firstValue));
```

Chiamare il metodo `handleClick` su `MyComponent` dovrebbe impostare la proprietà name nello stato uguale a `React Rocks!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'Before' });
    return waitForIt(() => mockedComponent.state('name'));
  };
  const second = () => {
    mockedComponent.instance().handleClick();
    return waitForIt(() => mockedComponent.state('name'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Before' && secondValue === 'React Rocks!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick={this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'Initial State'
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
     // Change code below this line
    this.setState({
      name: 'React Rocks!'
    });
    // Change code above this line
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.name}</h1>
      </div>
    );
  }
};
```
