---
id: 5a24c314108439a4d403616d
title: Use PropTypes to Define the Props You Expect
challengeType: 6
forumTopicId: 301419
dashedName: use-proptypes-to-define-the-props-you-expect
---

# --description--

React offre utili funzionalità di controllo dei tipi per verificare che i componenti ricevano props del tipo corretto. Per esempio, la tua applicazione fa una chiamata API per recuperare dati che ti aspetti siano in un array, che poi vengono passati a un componente come prop. Puoi impostare `propTypes` sul tuo componente per richiedere che i dati siano di tipo `array`. Questo genererà un avviso utile quando i dati sono di un tipo diverso.

È considerata una buona pratica impostare `propTypes` quando conosci in anticipo il tipo di una prop. Puoi definire una proprietà `propTypes` per un componente nello stesso modo in cui hai definito `defaultProps`. Così facendo, verificherai che le props con una certa chiave siano presenti e abbiano un tipo specifico. Ecco un esempio per richiedere il tipo `function` per una prop chiamata `handleClick`:

```js
MyComponent.propTypes = { handleClick: PropTypes.func.isRequired }
```

Nell'esempio sopra, la parte `PropTypes.func` verifica che `handleClick` sia una funzione. Aggiungere `isRequired` dice a React che `handleClick` è una proprietà obbligatoria per quel componente. Vedrai un avviso se quella prop non viene fornita. Nota anche che `func` rappresenta `function`. Tra i sette tipi primitivi di JavaScript, `function` e `boolean` (scritti come `bool`) sono gli unici due con una grafia insolita. Oltre ai tipi primitivi, ci sono altri tipi disponibili. Per esempio, puoi verificare che una prop sia un elemento React. Consulta la <a href="https://reactjs.org/docs/typechecking-with-proptypes.html#proptypes" target="_blank" rel="noopener noreferrer nofollow">documentazione</a> per tutte le opzioni.

**Nota:** A partire da React v15.5.0, `PropTypes` si importa indipendentemente da React, così: `import PropTypes from 'prop-types';`

# --instructions--

Definisci `propTypes` per il componente `Items` per richiedere `quantity` come prop e verificare che sia di tipo `number`.

# --hints--

Il componente `ShoppingCart` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Il componente `Items` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Il componente `Items` dovrebbe includere un controllo `propTypes` per richiedere un valore per `quantity` e assicurarsi che il suo valore sia un numero.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(noWhiteSpace.includes('quantity:PropTypes.number.isRequired') && noWhiteSpace.includes('Items.propTypes='));
```

# --before-all--

```js
var PropTypes = {
  number: { isRequired: true }
};
```

# --seed--

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line

// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
};

// Change code below this line
Items.propTypes = {
  quantity: PropTypes.number.isRequired
};
// Change code above this line

Items.defaultProps = {
  quantity: 0
};

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <Items />
  }
};
```
