---
id: 5a24c314108439a4d403618b
title: Give Sibling Elements a Unique Key Attribute
challengeType: 6
forumTopicId: 301394
dashedName: give-sibling-elements-a-unique-key-attribute
---

# --description--

L'ultima sfida ha mostrato come il metodo `map` venga usato per renderizzare dinamicamente un numero di elementi basandosi sull'input dell'utente. Tuttavia, in quell'esempio mancava un dettaglio importante. Quando crei un array di elementi, ciascuno deve avere un attributo `key` impostato su un valore unico. React usa queste key per tenere traccia di quali elementi sono stati aggiunti, modificati o rimossi. Questo aiuta a rendere il processo di re-rendering più efficiente quando la lista viene modificata in qualsiasi modo.

**Nota:** Le key devono essere uniche solo tra elementi fratelli, non è necessario che siano uniche globalmente nella tua applicazione.

# --instructions--

L'editor del codice contiene un array con alcuni framework front-end e un componente funzionale senza stato chiamato `Frameworks()`. `Frameworks()` deve mappare l'array in una lista non ordinata, proprio come nell'ultima sfida. Completa la scrittura della callback `map` per restituire un elemento `li` per ogni framework nell'array `frontEndFrameworks`. Questa volta, assicurati di dare a ogni `li` un attributo `key` impostato su un valore unico. Gli elementi `li` devono inoltre contenere il testo proveniente da `frontEndFrameworks`.

Normalmente, vuoi che la key sia qualcosa che identifica in modo univoco l'elemento che viene renderizzato. Come ultima risorsa si può usare l'indice dell'array, ma in genere dovresti cercare di usare un identificatore unico.

# --hints--

Il componente `Frameworks` dovrebbe esistere e renderizzare sulla pagina.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('Frameworks').length === 1
);
```

`Frameworks` dovrebbe renderizzare un elemento `h1`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('h1').length === 1);
```

`Frameworks` dovrebbe renderizzare un elemento `ul`.

```js
assert(Enzyme.mount(React.createElement(Frameworks)).find('ul').length === 1);
```

Il tag `ul` dovrebbe renderizzare 6 elementi figli `li`.

```js
assert(
  Enzyme.mount(React.createElement(Frameworks)).find('ul').children().length ===
    6 &&
    Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .childAt(0)
      .name() === 'li' &&
    Enzyme.mount(React.createElement(Frameworks)).find('li').length === 6
);
```

Ogni elemento della lista dovrebbe avere un attributo `key` unico.

```js
assert(
  (() => {
    const ul = Enzyme.mount(React.createElement(Frameworks)).find('ul');
    const keys = new Set([
      ul.childAt(0).key(),
      ul.childAt(1).key(),
      ul.childAt(2).key(),
      ul.childAt(3).key(),
      ul.childAt(4).key(),
      ul.childAt(5).key()
    ]);
    return keys.size === 6;
  })()
);
```

Ogni elemento della lista dovrebbe contenere testo proveniente da `frontEndFrameworks`.

```js
assert(
  (() => {
    const li = Enzyme.mount(React.createElement(Frameworks))
      .find('ul')
      .children();
    return [...Array(5)].every((_, i) =>
      frontEndFrameworks.includes(li.at(i).text())
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = null; // Change this line
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```

# --solutions--

```jsx
const frontEndFrameworks = [
  'React',
  'Angular',
  'Ember',
  'Knockout',
  'Backbone',
  'Vue'
];

function Frameworks() {
  const renderFrameworks = frontEndFrameworks.map((fw, i) => <li key={i}>{fw}</li>);
  return (
    <div>
      <h1>Popular Front-End JavaScript Frameworks</h1>
      <ul>
        {renderFrameworks}
      </ul>
    </div>
  );
};
```
