---
id: 5a24c314108439a4d4036180
title: ottimizza i re-render con shouldComponentUpdate
challengeType: 6
forumTopicId: 301398
dashedName: optimize-re-renders-with-shouldcomponentupdate
---

# --description--

Finora, se un componente riceve nuove `state` o nuove `props`, si ri-renderizza insieme a tutti i suoi figli. Di solito va bene così. Ma React fornisce un metodo del ciclo di vita che puoi chiamare quando i componenti figli ricevono nuove `state` o `props`, per dichiarare specificamente se i componenti devono aggiornarsi o no. Il metodo è `shouldComponentUpdate()`, e prende `nextProps` e `nextState` come parametri.

Questo metodo è un modo utile per ottimizzare le prestazioni. Per esempio, il comportamento predefinito è che il tuo componente si ri-renderizza quando riceve nuove `props`, anche se le `props` non sono cambiate. Puoi usare `shouldComponentUpdate()` per evitare questo confrontando le `props`. Il metodo deve restituire un valore `boolean` che dice a React se aggiornare o meno il componente. Puoi confrontare le props correnti (`this.props`) con le props successive (`nextProps`) per decidere se aggiornare o no, e restituire `true` o `false` di conseguenza.

# --instructions--

Il metodo `shouldComponentUpdate()` è aggiunto in un componente chiamato `OnlyEvens`. Attualmente, questo metodo restituisce `true` quindi `OnlyEvens` si ri-renderizza ogni volta che riceve nuove `props`. Modifica il metodo in modo che `OnlyEvens` si aggiorni solo se la `value` delle sue nuove props è pari. Clicca il pulsante `Add` e osserva l’ordine degli eventi nella console del browser mentre i metodi del ciclo di vita vengono attivati.

# --hints--

Il componente `Controller` dovrebbe renderizzare il componente `OnlyEvens` come figlio.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(Controller));
    return (
      mockedComponent.find('Controller').length === 1 &&
      mockedComponent.find('OnlyEvens').length === 1
    );
  })()
);
```

Il metodo `shouldComponentUpdate` dovrebbe essere definito nel componente `OnlyEvens`.

```js
assert(
  (() => {
    const child = React.createElement(OnlyEvens)
      .type.prototype.shouldComponentUpdate.toString()
      .replace(/s/g, '');
    return child !== 'undefined';
  })()
);
```

Il componente `OnlyEvens` dovrebbe restituire un tag `h1` che renderizza il valore di `this.props.value`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 1000 });
    return mockedComponent.find('h1').html();
  };
  const second = () => {
    mockedComponent.setState({ value: 10 });
    return mockedComponent.find('h1').html();
  };
  const firstValue = first();
  const secondValue = second();
  assert(firstValue === '<h1>1000</h1>' && secondValue === '<h1>10</h1>');
})();
```

`OnlyEvens` dovrebbe ri-renderizzarsi solo quando `nextProps.value` è pari.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(Controller));
  const first = () => {
    mockedComponent.setState({ value: 8 });
    return mockedComponent.find('h1').text();
  };
  const second = () => {
    mockedComponent.setState({ value: 7 });
    return mockedComponent.find('h1').text();
  };
  const third = () => {
    mockedComponent.setState({ value: 42 });
    return mockedComponent.find('h1').text();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(firstValue === '8' && secondValue === '8' && thirdValue === '42');
})();
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<Controller />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return true;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```

# --solutions--

```jsx
class OnlyEvens extends React.Component {
  constructor(props) {
    super(props);
  }
  shouldComponentUpdate(nextProps, nextState) {
    console.log('Should I update?');
    // Change code below this line
    return nextProps.value % 2 === 0;
    // Change code above this line
  }
  componentDidUpdate() {
    console.log('Component re-rendered.');
  }
  render() {
    return <h1>{this.props.value}</h1>;
  }
}

class Controller extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      value: 0
    };
    this.addValue = this.addValue.bind(this);
  }
  addValue() {
    this.setState(state => ({
      value: state.value + 1
    }));
  }
  render() {
    return (
      <div>
        <button onClick={this.addValue}>Add</button>
        <OnlyEvens value={this.state.value} />
      </div>
    );
  }
}
```
