---
id: 5a24c314108439a4d4036185
title: Usa && per una condizione più concisa
challengeType: 6
forumTopicId: 301413
dashedName: use--for-a-more-concise-conditional
---

# --description--

Le istruzioni `if/else` hanno funzionato nella sfida precedente, ma c’è un modo più conciso per ottenere lo stesso risultato. Immagina di dover tenere traccia di diverse condizioni in un componente e di voler renderizzare elementi diversi a seconda di ciascuna di queste condizioni. Se scrivi molte istruzioni `else if` per restituire interfacce leggermente diverse, potresti ripetere codice, lasciando spazio a errori. Invece, puoi usare l’operatore logico `&&` per eseguire una logica condizionale in modo più conciso. Questo è possibile perché vuoi verificare se una condizione è `true` e, se lo è, restituire del markup. Ecco un esempio:

```jsx
{condition && <p>markup</p>}
```

Se la condizione `condition` è `true`, il markup verrà restituito. Se la condizione è `false`, l’operazione restituirà immediatamente `false` dopo aver valutato `condition` e non restituirà nulla. Puoi includere queste istruzioni direttamente nel tuo JSX e concatenare più condizioni scrivendo `&&` dopo ciascuna. Questo ti permette di gestire logiche condizionali più complesse nel metodo `render()` senza ripetere molto codice.

# --instructions--

Risolvi di nuovo l’esempio precedente, in modo che `h1` venga renderizzato solo se `display` è `true`, ma usa l’operatore logico `&&` invece di un’istruzione `if/else`.

# --hints--

`MyComponent` dovrebbe esistere e renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length;
  })()
);
```

Quando `display` è impostato su `true`, dovrebbero renderizzare un `div`, `button` e `h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 2 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 1
  );
```

Quando `display` è impostato su `false`, dovrebbero renderizzare solo un `div` e `button`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    updated.find('div').length === 1 &&
      updated.find('div').children().length === 1 &&
      updated.find('button').length === 1 &&
      updated.find('h1').length === 0
  );
```

Il metodo render dovrebbe usare l’operatore logico `&&` per verificare la condizione di `this.state.display`.

```js
assert(code.includes('&&'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState(state => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         {this.state.display && <h1>Displayed!</h1>}
       </div>
    );
  }
};
```
