---
id: 5a24c314108439a4d4036168
title: Scrivi un componente React da zero
challengeType: 6
forumTopicId: 301424
dashedName: write-a-react-component-from-scratch
---

# --description--

Ora che hai imparato le basi di JSX e dei componenti React, è il momento di scrivere un componente da solo. I componenti React sono i mattoni fondamentali delle applicazioni React, quindi è importante acquisire molta familiarità nel scriverli. Ricorda, un componente React tipico è una `class` ES6 che estende `React.Component`. Ha un metodo render che restituisce HTML (da JSX) o `null`. Questa è la forma base di un componente React. Una volta che la capisci bene, sarai pronto per iniziare a costruire progetti React più complessi.

# --instructions--

Definisci una classe `MyComponent` che estende `React.Component`. Il suo metodo render dovrebbe restituire un `div` che contiene un tag `h1` con il testo: `My First React Component!` al suo interno. Usa esattamente questo testo, maiuscole, minuscole e punteggiatura sono importanti. Assicurati anche di chiamare il costruttore per il tuo componente.

Renderizza questo componente nel DOM usando `ReactDOM.render()`. C’è un `div` con `id='challenge-node'` a tua disposizione da usare.

# --hints--

Dovrebbe esserci un componente React chiamato `MyComponent`.

```js
assert(__helpers.removeWhiteSpace(code).includes('classMyComponentextendsReact.Component{'));
```

`MyComponent` dovrebbe contenere un tag `h1` con il testo `My First React Component!`. Maiuscole, minuscole e punteggiatura sono importanti.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('h1').text() === 'My First React Component!';
  })()
);
```

`MyComponent` dovrebbe renderizzare nel DOM.

```js
assert(document.getElementById('challenge-node').childNodes.length === 1);
```

`MyComponent` dovrebbe avere un costruttore che chiama `super` con `props`.

```js
assert(
  MyComponent.toString().includes('MyComponent(props)') &&
    MyComponent.toString().includes('_super.call(this, props)')
);
```

# --seed--

## --seed-contents--

```jsx
// Change code below this line
```

# --solutions--

```jsx
// Change code below this line
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>My First React Component!</h1>
      </div>
    );
  }
};

ReactDOM.render(<MyComponent />, document.getElementById('challenge-node'));
```
