---
id: 5a24c314108439a4d4036184
title: Renderizzare con una condizione if-else
challengeType: 6
forumTopicId: 301410
dashedName: render-with-an-if-else-condition
---

# --description--

Un altro modo di usare JavaScript per controllare la vista renderizzata è collegare gli elementi renderizzati a una condizione. Quando la condizione è vera, viene renderizzata una vista. Quando è falsa, viene renderizzata una vista diversa. Puoi farlo con una normale istruzione `if/else` nel metodo `render()` di un componente React.

# --instructions--

MyComponent contiene uno stato `boolean` che tiene traccia se vuoi mostrare o meno un elemento nell’interfaccia utente. Il `button` cambia lo stato di questo valore. Attualmente, renderizza sempre la stessa interfaccia. Riscrivi il metodo `render()` con un’istruzione `if/else` in modo che se `display` è `true`, ritorni il markup attuale. Altrimenti, ritorna il markup senza l’elemento `h1`.

**Nota:** Devi scrivere un’istruzione `if/else` per superare i test. L’uso dell’operatore ternario non passerà qui.

# --hints--

`MyComponent` dovrebbe esistere e renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('MyComponent').length === 1;
  })()
);
```

Quando `display` è impostato su `true`, dovrebbero renderizzare un `div`, `button` e `h1`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: true });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 2 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 1
  );
```

Quando `display` è impostato su `false`, dovrebbero renderizzare solo un `div` e `button`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const state_1 = () => {
    mockedComponent.setState({ display: false });
    return waitForIt(() => mockedComponent);
  };
  const updated = await state_1();
  assert(
    mockedComponent.find('div').length === 1 &&
      mockedComponent.find('div').children().length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('h1').length === 0
  );
```

Il metodo render dovrebbe usare un’istruzione `if/else` per controllare la condizione di `this.state.display`.

```js
assert(code.includes('if') && code.includes('else'));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
    this.toggleDisplay = this.toggleDisplay.bind(this);
  }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line

    return (
       <div>
         <button onClick={this.toggleDisplay}>Toggle Display</button>
         <h1>Displayed!</h1>
       </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      display: true
    }
 this.toggleDisplay = this.toggleDisplay.bind(this);
 }
  toggleDisplay() {
    this.setState((state) => ({
      display: !state.display
    }));
  }
  render() {
    // Change code below this line
    if (this.state.display) {
      return (
         <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
           <h1>Displayed!</h1>
         </div>
      );
    } else {
      return (
        <div>
           <button onClick={this.toggleDisplay}>Toggle Display</button>
         </div>
      );
    }
  }
};
```
