---
id: 5a24c314108439a4d4036172
title: Render State in the User Interface Another Way
challengeType: 6
forumTopicId: 301408
dashedName: render-state-in-the-user-interface-another-way
---

# --description--

C'è un altro modo per consultare `state` in un componente. Nel metodo `render()`, prima dell'istruzione `return`, puoi scrivere JavaScript direttamente. Per esempio, potresti dichiarare funzioni, consultare dati da `state` o `props`, eseguire calcoli su questi dati e così via. Poi, puoi assegnare qualsiasi dato a variabili, a cui hai accesso nell'istruzione `return`.

# --instructions--

Nel metodo di rendering `MyComponent`, definisci una `const` chiamata `name` e assegnale il valore name nello `state` dello stato del componente. Poiché puoi scrivere JavaScript direttamente in questa parte del codice, non devi racchiudere questo riferimento tra parentesi graffe.

Poi, nell'istruzione return, renderizza questo valore in un tag `h1` usando la variabile `name`. Ricorda, devi usare la sintassi JSX (parentesi graffe per JavaScript) nell'istruzione return.

# --hints--

`MyComponent` dovrebbe avere una chiave `name` con valore `freeCodeCamp` archiviato nel suo stato.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('name') ===
    'freeCodeCamp'
);
```

`MyComponent` dovrebbe renderizzare un elemento di intestazione `h1` racchiuso in un singolo `div`.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.mount(React.createElement(MyComponent)).html()
  )
);
```

Il tag `h1` renderizzato dovrebbe includere un riferimento a `{name}`.

```js
assert(/<h1>\n*\s*\{\s*name\s*\}\s*\n*<\/h1>/.test(code));
```

L’elemento intestazione `h1` renderizzato dovrebbe contenere testo renderizzato dallo stato del componente.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.html());
  };
  const firstValue = await first();
  assert(firstValue === '<div><h1>TestName</h1></div>');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line

    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'freeCodeCamp'
    }
  }
  render() {
    // Change code below this line
    const name = this.state.name;
    // Change code above this line
    return (
      <div>
        { /* Change code below this line */ }
        <h1>{name}</h1>
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
