---
id: 5a24c314108439a4d4036181
title: Introducing Inline Styles
challengeType: 6
forumTopicId: 301395
dashedName: introducing-inline-styles
---

# --description--

Ci sono altri concetti complessi che aggiungono potenti funzionalità al tuo codice React. Ma potresti chiederti come risolvere il problema più semplice di come definire lo stile per quegli elementi JSX che crei in React. Probabilmente sai che non sarà esattamente uguale a lavorare con HTML a causa di <a href="/learn/front-end-development-libraries/react/define-an-html-class-in-jsx" target="_blank" rel="noopener noreferrer nofollow">come si applicano le classi agli elementi JSX</a>.

Se importi gli stili da un foglio di stile, non è molto diverso. Applichi una classe al tuo elemento JSX usando l'attributo `className` e definisci gli stili per la classe nel foglio di stile. Un'altra opzione è definire stili in linea, molto comuni nello sviluppo ReactJS.

Applichi gli stili in linea agli elementi JSX in modo simile a come fai in HTML, ma con alcune differenze di JSX. Ecco un esempio di stile in linea in HTML:

```jsx
<div style="color: yellow; font-size: 16px">Mellow Yellow</div>
```

Gli elementi JSX usano l'attributo `style`, ma a causa di come JSX viene trascritto, non puoi impostare il valore a una `string`. Invece, lo imposti uguale a un `object` JavaScript. Ecco un esempio:

```jsx
<div style={{color: "yellow", fontSize: 16}}>Mellow Yellow</div>
```

Nota come usiamo camelCase per la proprietà `fontSize`? Questo perché React non accetta chiavi in kebab-case nell'oggetto style. React applicherà il nome corretto della proprietà per noi in HTML.

# --instructions--

Aggiungi un attributo `style` al `div` nell'editor di codice per dare al testo un colore rosso e una dimensione del carattere di `72px`.

Nota che puoi opzionalmente impostare la dimensione del carattere come numero, omettendo le unità `px`, oppure scriverla come `"72px"`.
# --hints--

Il componente dovrebbe renderizzare un elemento `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().type() === 'div';
  })()
);
```

L'elemento `div` dovrebbe avere un colore `red`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return mockedComponent.children().props().style.color === 'red';
  })()
);
```

L'elemento `div` dovrebbe avere una dimensione del carattere `72px`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Colorful));
    return (
      mockedComponent.children().props().style.fontSize === 72 ||
      mockedComponent.children().props().style.fontSize === '72' ||
      mockedComponent.children().props().style.fontSize === '72px'
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div>Big Red</div>
    );
  }
};
```

# --solutions--

```jsx
class Colorful extends React.Component {
  render() {
    return (
      <div style={{color: "red", fontSize: 72}}>Big Red</div>
    );
  }
};
```
