---
id: 5a24c314108439a4d403616f
title: "Revisione sull'uso delle props con componenti funzionali senza stato"
challengeType: 6
forumTopicId: 301411
dashedName: review-using-props-with-stateless-functional-components
---

# --description--

Tranne che per l'ultima sfida, hai passato props a componenti funzionali senza stato. Questi componenti si comportano come funzioni pure. Accettano props come input e restituiscono la stessa vista ogni volta che ricevono le stesse props. Potresti chiederti cos'è lo stato, e la sfida successiva lo tratterà in modo più dettagliato. Prima di questo, ecco una revisione della terminologia per i componenti.

Un *componente funzionale senza stato* è qualsiasi funzione che scrivi che accetta props e restituisce JSX. Un *componente senza stato*, invece, è una classe che estende `React.Component`, ma non usa uno stato interno (trattato nella sfida successiva). Infine, un *componente con stato* è un componente di classe che mantiene il proprio stato interno. Potresti vedere i componenti con stato indicati semplicemente come componenti o componenti React.

Un pattern comune è cercare di minimizzare i componenti con stato e creare componenti funzionali senza stato ovunque sia possibile. Questo aiuta a contenere la gestione dello stato in un'area specifica della tua applicazione. A sua volta, questo migliora lo sviluppo e la manutenzione della tua app rendendo più facile seguire come le modifiche allo stato influenzano il suo comportamento.

# --instructions--

L'editor del codice ha un componente `CampSite` che renderizza un componente `Camper` come figlio. Definisci il componente `Camper` e assegna delle props predefinite di `{ name: 'CamperBot' }`. All'interno del componente `Camper`, renderizza qualsiasi codice tu voglia, ma assicurati di avere un elemento `p` che includa solo il valore `name` passato come `prop`. Infine, definisci `propTypes` sul componente `Camper` per richiedere che `name` sia fornito come prop e verifica che sia di tipo `string`.

# --hints--

Il componente `CampSite` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('CampSite').length === 1;
  })()
);
```

Il componente `Camper` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return mockedComponent.find('Camper').length === 1;
  })()
);
```

Il componente `Camper` dovrebbe includere props predefinite che assegnano la stringa `CamperBot` alla chiave `name`.

```js
assert(
  /Camper.defaultProps={name:(['"`])CamperBot\1,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Il componente `Camper` dovrebbe includere prop types che richiedono che la prop `name` sia di tipo `string`.

```js
assert(
  /Camper.propTypes={name:PropTypes.string.isRequired,?}/.test(
    __helpers.removeWhiteSpace(code)
  )
);
```

Il componente `Camper` dovrebbe contenere un elemento `p` con solo il testo della prop `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(CampSite));
    return (
      mockedComponent.find('p').text() ===
      mockedComponent.find('Camper').props().name
    );
  })()
);
```

# --seed--

## --before-user-code--

```jsx
var PropTypes = {
   string: { isRequired: true }
};
```

## --after-user-code--

```jsx
ReactDOM.render(<CampSite />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line
```

# --solutions--

```jsx
class CampSite extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <Camper/>
      </div>
    );
  }
};
// Change code below this line

const Camper = (props) => {
   return (
     <div>
       <p>{props.name}</p>
     </div>
   );
};

Camper.propTypes = {
  name: PropTypes.string.isRequired
};

Camper.defaultProps = {
  name: 'CamperBot'
};
```
