---
id: 5a24c314108439a4d4036169
title: Pass Props to a Stateless Functional Component
challengeType: 6
forumTopicId: 301402
dashedName: pass-props-to-a-stateless-functional-component
---

# --description--

Le sfide precedenti hanno coperto molto sulla creazione e composizione di elementi JSX, componenti funzionali e componenti di classe in stile ES6 in React. Con queste basi, è ora di guardare un'altra funzionalità molto comune in React: le **props**. In React, puoi passare props, o proprietà, ai componenti figli. Supponiamo che tu abbia un componente `App` che renderizza un componente figlio chiamato `Welcome`, che è un componente funzionale senza stato. Puoi passare a `Welcome` una proprietà `user` scrivendo:

```jsx
<App>
  <Welcome user='Mark' />
</App>
```

Usi **attributi HTML personalizzati** creati da te e supportati da React per passarli al componente. In questo caso, la proprietà creata `user` viene passata al componente `Welcome`. Poiché `Welcome` è un componente funzionale senza stato, ha accesso a questo valore così:

```jsx
const Welcome = (props) => <h1>Hello, {props.user}!</h1>
```

È consuetudine chiamare questo valore `props` e quando si lavora con componenti funzionali senza stato, lo consideri fondamentalmente come un argomento di una funzione che restituisce JSX. Puoi consultare il valore dell'argomento nel corpo della funzione. Con i componenti di classe, vedrai che è un po' diverso.

# --instructions--

Ci sono `Calendar` e `CurrentDate` componenti nell'editor di codice. Quando renderizzi `CurrentDate` dal componente `Calendar`, passa una proprietà `date` assegnata alla data corrente dall'oggetto `Date` di JavaScript. Poi consulta questa `prop` nel componente `CurrentDate`, mostrando il suo valore all'interno dei tag `p`. Nota che per far sì che i valori `prop` vengano valutati come JavaScript, devono essere racchiusi tra parentesi graffe, per esempio `date={Date()}`.

# --hints--

Il componente `Calendar` dovrebbe restituire un singolo elemento `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Il secondo figlio del componente `Calendar` dovrebbe essere il componente `CurrentDate`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).name() === 'CurrentDate';
  })()
);
```

Il componente `CurrentDate` dovrebbe avere una prop chiamata `date`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    return mockedComponent.children().childAt(1).props().date;
  })()
);
```

La prop `date` del componente `CurrentDate` dovrebbe contenere una stringa di testo.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(Calendar));
    const prop = mockedComponent.children().childAt(1).props().date;
    return typeof prop === 'string' && prop.length > 0;
  })()
);
```

La prop `date` dovrebbe essere generata chiamando `Date()`

```js
assert(/<CurrentDatedate={Date\(\)}\/>/.test(__helpers.removeWhiteSpace(code)));
```

Il componente `CurrentDate` dovrebbe renderizzare il valore della prop `date` nel tag `p`.

```js
let date = 'dummy date';
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(CurrentDate, { date })
    );
    return mockedComponent.find('p').html().includes(date);
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: </p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const CurrentDate = (props) => {
  return (
    <div>
      { /* Change code below this line */ }
      <p>The current date is: {props.date}</p>
      { /* Change code above this line */ }
    </div>
  );
};

class Calendar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>What date is it?</h3>
        { /* Change code below this line */ }
        <CurrentDate date={Date()} />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
