---
id: 5a24c314108439a4d4036163
title: Create a React Component
challengeType: 6
forumTopicId: 301386
dashedName: create-a-react-component
---

# --description--

L'altro modo per definire un componente React è con la sintassi ES6 `class`. Nell'esempio seguente, `Kitten` estende `React.Component`:

```jsx
class Kitten extends React.Component {
  constructor(props) {
    super(props);
  }

  render() {
    return (
      <h1>Hi</h1>
    );
  }
}
```

Questo crea una classe ES6 `Kitten` che estende la classe `React.Component`. Quindi la classe `Kitten` ha ora accesso a molte funzionalità React utili, come lo stato locale e i `lifecycle hooks`. Non preoccuparti se non conosci ancora questi termini, saranno trattati in modo più approfondito nelle sfide successive. Nota anche che la classe `Kitten` ha un `constructor` definito al suo interno che chiama `super()`. Usa `super()` per chiamare il costruttore della classe genitore, in questo caso `React.Component`. Il costruttore è un metodo speciale usato durante l'inizializzazione degli oggetti creati con la parola chiave `class`. È buona pratica chiamare il costruttore di un componente con `super`, e passare `props` a entrambi. Questo assicura che il componente venga inizializzato correttamente. Per ora, sappi che è standard includere questo codice. Presto vedrai altri usi per il costruttore così come `props`.

# --instructions--

`MyComponent` è definito nell'editor di codice usando la sintassi class. Completa la scrittura del metodo `render` in modo che ritorni un elemento `div` che contenga un `h1` con il testo `Hello React!`.

# --hints--

Il componente React dovrebbe restituire un elemento `div`.

```js
assert(Enzyme.shallow(React.createElement(MyComponent)).type() === 'div');
```

L'elemento `div` restituito dovrebbe renderizzare al suo interno un elemento di intestazione `h1`.

```js
assert(
  /<div><h1>.*<\/h1><\/div>/.test(
    Enzyme.shallow(React.createElement(MyComponent)).html()
  )
);
```

L'elemento di intestazione `h1` dovrebbe contenere la stringa `Hello React!`.

```js
assert(
  Enzyme.shallow(React.createElement(MyComponent)).html() ===
    '<div><h1>Hello React!</h1></div>'
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line



    // Change code above this line
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    // Change code below this line
    return (
      <div>
        <h1>Hello React!</h1>
      </div>
    );
    // Change code above this line
  }
};
```
