---
id: 5a24c314108439a4d403617e
title: Add Event Listeners
challengeType: 6
forumTopicId: 301377
dashedName: add-event-listeners
---

# --description--

Il metodo `componentDidMount()` è anche il posto migliore per allegare qualsiasi gestore di eventi che devi aggiungere per funzionalità specifiche. React fornisce un sistema di eventi sintetici che avvolge il sistema di eventi nativo presente nei browser. Questo significa che il sistema di eventi sintetici si comporta esattamente allo stesso modo indipendentemente dal browser dell'utente — anche se gli eventi nativi possono comportarsi in modo diverso tra browser diversi.

Hai già usato alcuni di questi gestori di eventi sintetici come `onClick()`. Il sistema di eventi sintetici di React è ottimo da usare per la maggior parte delle interazioni che gestirai sugli elementi DOM. Tuttavia, se vuoi allegare un gestore di eventi agli oggetti document o window, devi farlo direttamente.

# --instructions--

Allega un gestore di eventi nel metodo `componentDidMount()` per gli eventi `keydown` e fai in modo che questi eventi attivino la callback `handleKeyPress()`. Puoi usare `document.addEventListener()` che prende l'evento (tra virgolette) come primo argomento e la callback come secondo argomento.

Poi, in `componentWillUnmount()`, rimuovi lo stesso gestore di eventi. Puoi passare gli stessi argomenti a `document.removeEventListener()`. È buona pratica usare questo metodo del ciclo di vita per fare qualsiasi pulizia sui componenti React prima che vengano smontati e distrutti. Rimuovere i gestori di eventi è un esempio di questo tipo di azione di pulizia.

# --hints--

`MyComponent` dovrebbe renderizzare un elemento `div` che avvolge un tag `h1`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').children().find('h1').length === 1;
  })()
);
```

Un gestore `keydown` dovrebbe essere allegato al document in `componentDidMount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const didMountString = mockedComponent
      .instance()
      .componentDidMount.toString();
    return new RegExp(
      'document.addEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(didMountString);
  })()
);
```

Il gestore `keydown` dovrebbe essere rimosso dal document in `componentWillUnmount`.

```js
assert(
  (() => {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    const willUnmountString = mockedComponent
      .instance()
      .componentWillUnmount.toString();
    return new RegExp(
      'document.removeEventListener(.|\n|\r)+keydown(.|\n|\r)+this.handleKeyPress'
    ).test(willUnmountString);
  })()
);
```

Una volta che il componente è montato, premere `enter` dovrebbe aggiornare il suo stato e il tag `h1` renderizzato.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const beforeState = mockedComponent.state('message');
  const beforeText = mockedComponent.find('h1').text();
  const pressEnterKey = () => {
    mockedComponent.instance().handleKeyPress({ keyCode: 13 });
    return waitForIt(() => {
      mockedComponent.update();
      return {
        state: mockedComponent.state('message'),
        text: mockedComponent.find('h1').text()
      };
    });
  };
  const afterKeyPress = await pressEnterKey();
  assert(
    beforeState !== afterKeyPress.state && beforeText !== afterKeyPress.text
  );
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleEnter = this.handleEnter.bind(this);
    this.handleKeyPress = this.handleKeyPress.bind(this);
  }
  // Change code below this line
  componentDidMount() {

  }
  componentWillUnmount() {

  }
  // Change code above this line
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      message: ''
    };
    this.handleKeyPress = this.handleKeyPress.bind(this);
    this.handleEnter = this.handleEnter.bind(this);  }
  componentDidMount() {
    // Change code below this line
    document.addEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  componentWillUnmount() {
    // Change code below this line
    document.removeEventListener('keydown', this.handleKeyPress);
    // Change code above this line
  }
  handleEnter() {
    this.setState((state) => ({
      message: state.message + 'You pressed the enter key! '
    }));
  }
  handleKeyPress(event) {
    if (event.keyCode === 13) {
      this.handleEnter();
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.message}</h1>
      </div>
    );
  }
};
```
