---
id: 5a24c314108439a4d4036174
title: Bind 'this' to a Class Method
challengeType: 6
forumTopicId: 301379
dashedName: bind-this-to-a-class-method
---

# --description--

Oltre a impostare e aggiornare `state`, puoi anche definire metodi per la tua classe componente. Un metodo di classe di solito deve usare la parola chiave `this` per poter consultare le proprietà della classe (come `state` e `props`) all'interno dell'ambito del metodo. Ci sono diversi modi per permettere ai metodi della tua classe di consultare `this`.

Un modo comune è collegare esplicitamente `this` nel costruttore in modo che `this` venga collegato ai metodi della classe quando il componente viene inizializzato. Avrai notato che nell'ultima sfida è stato usato `this.handleClick = this.handleClick.bind(this)` per il metodo `handleClick` nel costruttore. Poi, quando chiami una funzione come `this.setState()` all'interno del metodo della tua classe, `this` si riferisce alla classe e non sarà `undefined`.

**Nota:** La parola chiave `this` è uno degli aspetti più confusi di JavaScript ma ha un ruolo importante in React. Anche se il suo comportamento qui è del tutto normale, queste lezioni non sono il posto giusto per una revisione approfondita di `this`, quindi ti invitiamo a consultare altre lezioni se quanto sopra ti risulta confuso!

# --instructions--

L'editor del codice ha un componente con uno `state` che tiene traccia del testo. Ha anche un metodo che ti permette di impostare il testo su `You clicked!`. Tuttavia, il metodo non funziona perché usa la parola chiave `this` che è indefinita. Sistemalo collegando esplicitamente `this` al metodo `handleClick()` nel costruttore del componente.

Poi, aggiungi un gestore di clic all'elemento `button` nel metodo render. Dovrebbe attivare il metodo `handleClick()` quando il pulsante riceve un evento di clic. Ricorda che il metodo che passi al gestore `onClick` deve essere racchiuso tra parentesi graffe perché deve essere interpretato direttamente come JavaScript.

Una volta completati i passaggi sopra, dovresti poter cliccare il pulsante e vedere `You clicked!`.

# --hints--

`MyComponent` dovrebbe restituire un elemento `div` che racchiude due elementi, un pulsante e un elemento `h1`, in quest'ordine.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).find('div').length === 1 &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(0)
      .type() === 'button' &&
    Enzyme.mount(React.createElement(MyComponent))
      .find('div')
      .childAt(1)
      .type() === 'h1'
);
```

Lo stato di `MyComponent` dovrebbe inizializzarsi con la coppia chiave-valore `{ text: "Hello" }`.

```js
assert(
  Enzyme.mount(React.createElement(MyComponent)).state('text') === 'Hello'
);
```

Cliccando sull'elemento `button` dovrebbe eseguire il metodo `handleClick` e impostare lo stato `text` su `You clicked!`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ text: 'Hello' });
    return waitForIt(() => mockedComponent.state('text'));
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return waitForIt(() => mockedComponent.state('text'));
  };
  const firstValue = await first();
  const secondValue = await second();
  assert(firstValue === 'Hello' && secondValue === 'You clicked!');
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    // Change code below this line

    // Change code above this line
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        { /* Change code below this line */ }
        <button>Click Me</button>
        { /* Change code above this line */ }
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      text: "Hello"
    };
    this.handleClick = this.handleClick.bind(this);
  }
  handleClick() {
    this.setState({
      text: "You clicked!"
    });
  }
  render() {
    return (
      <div>
        <button onClick = {this.handleClick}>Click Me</button>
        <h1>{this.state.text}</h1>
      </div>
    );
  }
};
```
