---
id: 5a24c314108439a4d403617a
title: Pass State as Props to Child Components
challengeType: 6
forumTopicId: 301403
dashedName: pass-state-as-props-to-child-components
---

# --description--

Hai visto molti esempi che passavano props agli elementi JSX figli e ai componenti React figli nelle sfide precedenti. Ti starai chiedendo da dove provengano quelle props. Un pattern comune è avere un componente con stato che contiene il `state` importante per la tua app, che poi renderizza componenti figli. Vuoi che questi componenti abbiano accesso ad alcune parti di quello `state`, che vengono passate come props.

Per esempio, magari hai un componente `App` che renderizza un `Navbar`, tra gli altri componenti. Nel tuo `App`, hai un `state` che contiene molte informazioni sull’utente, ma il `Navbar` ha bisogno solo di accedere al nome utente per poterlo mostrare. Passi quella parte di `state` al componente `Navbar` come prop.

Questo pattern illustra alcuni paradigmi importanti in React. Il primo è il *flusso dati unidirezionale*. Lo stato scorre in una sola direzione lungo l’albero dei componenti della tua applicazione, dal componente genitore con stato ai componenti figli. I componenti figli ricevono solo i dati di stato di cui hanno bisogno. Il secondo è che app complesse con stato possono essere suddivise in pochi, o magari un solo, componente con stato. Il resto dei componenti riceve semplicemente lo stato dal genitore come props e renderizza un’interfaccia utente da quello stato. Si crea così una separazione in cui la gestione dello stato è gestita in una parte del codice e il rendering dell’interfaccia in un’altra. Questo principio di separare la logica dello stato dalla logica dell’interfaccia è uno dei principi chiave di React. Quando viene usato correttamente, rende molto più facile gestire il design di applicazioni complesse con stato.

# --instructions--

Il componente `MyApp` ha stato e renderizza un componente `Navbar` come figlio. Passa la proprietà `name` nel suo `state` al componente figlio, poi mostra il `name` nel tag `h1` che fa parte del metodo di render `Navbar`. `name` dovrebbe apparire dopo il testo `Hello, my name is:`.

# --hints--

Il componente `MyApp` dovrebbe renderizzare un componente `Navbar` all’interno.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return (
      mockedComponent.find('MyApp').length === 1 &&
      mockedComponent.find('Navbar').length === 1
    );
  })()
);
```

Il componente `Navbar` dovrebbe ricevere la proprietà di stato `MyApp` `name` come props.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').props());
  };
  const navProps = await setState();
  assert(navProps.name === 'TestName');
```

L’elemento `h1` in `Navbar` dovrebbe renderizzare la prop `name`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const navH1Before = mockedComponent.find('Navbar').find('h1').text();
  const setState = () => {
    mockedComponent.setState({ name: 'TestName' });
    return waitForIt(() => mockedComponent.find('Navbar').find('h1').text());
  };
  const navH1After = await setState();
  assert(new RegExp('TestName').test(navH1After) && navH1After !== navH1Before);
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         {/* Change code below this line */}
         <Navbar />
         {/* Change code above this line */}
       </div>
    );
  }
};

class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      {/* Change code below this line */}
      <h1>Hello, my name is: </h1>
      {/* Change code above this line */}
    </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      name: 'CamperBot'
    }
  }
  render() {
    return (
       <div>
         <Navbar name={this.state.name}/>
       </div>
    );
  }
};
class Navbar extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
    <div>
      <h1>Hello, my name is: {this.props.name}</h1>
    </div>
    );
  }
};
```
