---
id: 5a24c314108439a4d4036164
title: Create a Component with Composition
challengeType: 6
forumTopicId: 301383
dashedName: create-a-component-with-composition
---

# --description--

Ora vedremo come comporre insieme più componenti React. Immagina di costruire un'app e di aver creato tre componenti: un `Navbar`, `Dashboard` e `Footer`.

Per comporre questi componenti insieme, potresti creare un componente *genitore* `App` che renderizza ciascuno di questi tre componenti come *figli*. Per renderizzare un componente come figlio in un componente React, includi il nome del componente scritto come un tag HTML personalizzato nel JSX. Per esempio, nel metodo `render` potresti scrivere:

```jsx
return (
 <App>
  <Navbar />
  <Dashboard />
  <Footer />
 </App>
)
```

Quando React incontra un tag HTML personalizzato che fa riferimento a un altro componente (un nome di componente racchiuso in `< />` come in questo esempio), renderizza il markup per quel componente nella posizione del tag. Questo illustra la relazione genitore/figlio tra il componente `App` e i componenti `Navbar`, `Dashboard` e `Footer`.

# --instructions--

Nell'editor di codice c'è un semplice componente funzionale chiamato `ChildComponent` e un componente classe chiamato `ParentComponent`. Componili insieme renderizzando `ChildComponent` all'interno di `ParentComponent`. Assicurati di chiudere il tag `ChildComponent` con una barra obliqua.

**Nota:** `ChildComponent` è definito con una funzione freccia ES6 perché questa è una pratica molto comune quando si usa React.

# --hints--

Il componente React dovrebbe restituire un singolo elemento `div`.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.type() === 'div';
  })()
);
```

Il componente dovrebbe restituire due elementi annidati.

```js
assert(
  (function () {
    var shallowRender = Enzyme.shallow(React.createElement(ParentComponent));
    return shallowRender.children().length === 2;
  })()
);
```

Il componente dovrebbe restituire `ChildComponent` come suo secondo figlio.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ParentComponent));
    return (
      mockedComponent.find('ParentComponent').find('ChildComponent').length ===
      1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }


        { /* Change code above this line */ }
      </div>
    );
  }
};
```

# --solutions--

```jsx
const ChildComponent = () => {
  return (
    <div>
      <p>I am the child</p>
    </div>
  );
};

class ParentComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h1>I am the parent</h1>
        { /* Change code below this line */ }
        <ChildComponent />
        { /* Change code above this line */ }
      </div>
    );
  }
};
```
