---
id: 5a24c314108439a4d403617b
title: Pass a Callback as Props
challengeType: 6
forumTopicId: 301400
dashedName: pass-a-callback-as-props
---

# --description--

Puoi passare `state` come props ai componenti figli, ma non sei limitato a passare solo dati. Puoi anche passare funzioni gestore o qualsiasi metodo definito in un componente React a un componente figlio. Questo è il modo in cui permetti ai componenti figli di interagire con i loro componenti genitori. Passi i metodi a un figlio proprio come una prop normale. Gli assegni un nome e puoi consultare quel nome di metodo sotto `this.props` nel componente figlio.

# --instructions--

Ci sono tre componenti delineati nell'editor del codice. Il componente `MyApp` è il genitore che renderizzerà i componenti figli `GetInput` e `RenderInput`. Aggiungi il componente `GetInput` al metodo render in `MyApp`, quindi passagli una prop chiamata `input` assegnata a `inputValue` dal `MyApp` di `state`. Crea anche una prop chiamata `handleChange` e passale il gestore di input `handleChange`.

Poi, aggiungi `RenderInput` al metodo render in `MyApp`, quindi crea una prop chiamata `input` e passale `inputValue` da `state`. Una volta finito, potrai digitare nel campo `input` nel componente `GetInput`, che poi chiamerà il metodo gestore nel suo genitore tramite props. Questo aggiorna l'input nel `state` del genitore, che viene passato come props a entrambi i figli. Osserva come i dati fluiscono tra i componenti e come la singola fonte di verità rimane lo `state` del componente genitore. Ammettiamo che questo esempio sia un po' artificiale, ma dovrebbe servire a illustrare come dati e callback possono essere passati tra componenti React.

# --hints--

Il componente `MyApp` dovrebbe essere renderizzato.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('MyApp').length === 1;
  })()
);
```

Il componente `GetInput` dovrebbe essere renderizzato.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('GetInput').length === 1;
  })()
);
```

Il componente `RenderInput` dovrebbe essere renderizzato.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyApp));
    return mockedComponent.find('RenderInput').length === 1;
  })()
);
```

Il componente `GetInput` dovrebbe ricevere la proprietà di stato `MyApp` `inputValue` come props e contenere un elemento `input` che modifica lo stato `MyApp`.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: '' });
    return waitForIt(() => mockedComponent.state());
  };
  const state_2 = () => {
    mockedComponent
      .find('input')
      .simulate('change', { target: { value: 'TestInput' } });
    return waitForIt(() => mockedComponent.state());
  };
  const updated_1 = await state_1();
  const updated_2 = await state_2();
  assert(updated_1.inputValue === '' && updated_2.inputValue === 'TestInput');
```

Il componente `RenderInput` dovrebbe ricevere la proprietà di stato `MyApp` `inputValue` come props.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 250));
  const mockedComponent = Enzyme.mount(React.createElement(MyApp));
  const state_1 = () => {
    mockedComponent.setState({ inputValue: 'TestName' });
    return waitForIt(() => mockedComponent);
  };
  const updated_1 = await state_1();
  assert(updated_1.find('p').text().includes('TestName'));
```

# --seed--

## --seed-contents--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
        { /* Change code below this line */ }

        { /* Change code above this line */ }
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class MyApp extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      inputValue: ''
    }
  this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({
      inputValue: event.target.value
    });
  }
  render() {
    return (
       <div>
         <GetInput
           input={this.state.inputValue}
           handleChange={this.handleChange}/>
         <RenderInput
           input={this.state.inputValue}/>
       </div>
    );
  }
};

class GetInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Get Input:</h3>
        <input
          value={this.props.input}
          onChange={this.props.handleChange}/>
      </div>
    );
  }
};

class RenderInput extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return (
      <div>
        <h3>Input Render:</h3>
        <p>{this.props.input}</p>
      </div>
    );
  }
};
```
