---
id: 5a24c314108439a4d4036176
title: Usa lo stato per attivare un elemento
challengeType: 6
forumTopicId: 301421
dashedName: use-state-to-toggle-an-element
---

# --description--

A volte potresti aver bisogno di conoscere lo stato precedente quando aggiorni lo stato. Tuttavia, gli aggiornamenti di stato possono essere asincroni: questo significa che React può raggruppare più chiamate `setState()` in un unico aggiornamento. Questo implica che non puoi fare affidamento sul valore precedente di `this.state` o `this.props` quando calcoli il valore successivo. Quindi, non dovresti usare codice come questo:

```jsx
this.setState({
  counter: this.state.counter + this.props.increment
});
```

Invece, dovresti passare a `setState` una funzione che ti permette di consultare stato e props. Usare una funzione con `setState` garantisce che stai lavorando con i valori più aggiornati di stato e props. Questo significa che quanto sopra dovrebbe essere riscritto come:

```jsx
this.setState((state, props) => ({
  counter: state.counter + props.increment
}));
```

Puoi anche usare una forma senza `props` se ti serve solo il `state`:

```jsx
this.setState(state => ({
  counter: state.counter + 1
}));
```

Nota che devi racchiudere l'oggetto letterale tra parentesi tonde, altrimenti JavaScript lo interpreta come un blocco di codice.

# --instructions--

`MyComponent` ha una proprietà `visibility` che viene inizializzata a `false`. Il metodo renderizza restituisce una vista se il valore di `visibility` è true, e una vista diversa se è falso.

Attualmente non c’è modo di aggiornare la proprietà `visibility` nello stato `state` del componente. Il valore dovrebbe alternarsi tra true e false. C’è un gestore di clic sul pulsante che attiva un metodo di classe chiamato `toggleVisibility()`. Passa una funzione a `setState` per definire questo metodo in modo che la proprietà `state` di `visibility` si alterni al valore opposto quando il metodo viene chiamato. Se `visibility` è `false`, il metodo lo imposta a `true`, e viceversa.

Infine, clicca il pulsante per vedere il rendering condizionale del componente basato sul suo `state`.

**Suggerimento:** Non dimenticare di legare la parola chiave `this` al metodo nel `constructor`!

# --hints--

`MyComponent` dovrebbe restituire un elemento `div` che contiene un `button`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).find('div').find('button')
    .length,
  1
);
```

Lo stato di `MyComponent` dovrebbe inizializzarsi con una proprietà `visibility` impostata a `false`.

```js
assert.strictEqual(
  Enzyme.mount(React.createElement(MyComponent)).state('visibility'),
  false
);
```

Cliccando l’elemento pulsante dovrebbe alternare la proprietà `visibility` nello stato tra `true` e `false` e renderizzare condizionatamente l’elemento `h1`.

```js
(() => {
  const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
  const first = () => {
    mockedComponent.setState({ visibility: false });
    return mockedComponent.state('visibility');
  };
  const second = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const third = () => {
    mockedComponent.find('button').simulate('click');
    return mockedComponent.state('visibility') && mockedComponent.find('h1').exists();
  };
  const firstValue = first();
  const secondValue = second();
  const thirdValue = third();
  assert(!firstValue && secondValue && !thirdValue);
})();
```

Una funzione anonima dovrebbe essere passata a `setState`.

```js
const paramRegex = '[a-zA-Z$_]\\w*(,[a-zA-Z$_]\\w*)?';
assert(
  new RegExp(
    'this\\.setState\\((function\\(' +
      paramRegex +
      '\\){|([a-zA-Z$_]\\w*|\\(' +
      paramRegex +
      '\\))=>)'
  ).test(__helpers.removeWhiteSpace(code))
);
```

`this` non dovrebbe essere usato dentro `setState`

```js
assert(!/this\.setState\([^}]*this/.test(code));
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<MyComponent />, document.getElementById('root'));
```

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    // Change code below this line

    // Change code above this line
  }
  // Change code below this line

  // Change code above this line
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      visibility: false
    };
    this.toggleVisibility = this.toggleVisibility.bind(this);
  }
  toggleVisibility() {
    this.setState(state => ({
      visibility: !state.visibility
    }));
  }
  render() {
    if (this.state.visibility) {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
          <h1>Now you see me!</h1>
        </div>
      );
    } else {
      return (
        <div>
          <button onClick={this.toggleVisibility}>Click Me</button>
        </div>
      );
    }
  }
}
```
