---
id: 5a24c314108439a4d403616c
title: sovrascrivere le props predefinite
challengeType: 6
forumTopicId: 301399
dashedName: override-default-props
---

# --description--

La possibilità di impostare props predefinite è una funzionalità utile in React. Il modo per sovrascrivere le props predefinite è impostare esplicitamente i valori delle props per un componente.

# --instructions--

Il componente `ShoppingCart` ora renderizza un componente figlio `Items`. Questo componente `Items` ha una prop predefinita `quantity` impostata all'intero `0`. Sovrascrivi la prop predefinita passando un valore di `10` per `quantity`.

**Nota:** Ricorda che la sintassi per aggiungere una prop a un componente è simile a come si aggiungono gli attributi HTML. Tuttavia, dato che il valore per `quantity` è un intero, non va tra virgolette ma deve essere racchiuso tra parentesi graffe. Per esempio, `{100}`. Questa sintassi indica a JSX di interpretare direttamente come JavaScript il valore all'interno delle parentesi graffe.

# --hints--

Il componente `ShoppingCart` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Il componente `Items` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('Items').length === 1;
  })()
);
```

Il componente `Items` dovrebbe ricevere una prop `{ quantity: 10 }` passata dal componente `ShoppingCart`.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
      return (
        mockedComponent.find('Items').props().quantity == 10 &&
        code.replace(/ /g, '')
          .includes('<Itemsquantity={10}/>')
      );
    })()
  );
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<ShoppingCart />, document.getElementById('root'))
```

## --seed-contents--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items />
    { /* Change code above this line */ }
  }
};
```

# --solutions--

```jsx
const Items = (props) => {
  return <h1>Current Quantity of Items in Cart: {props.quantity}</h1>
}

Items.defaultProps = {
  quantity: 0
}

class ShoppingCart extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    { /* Change code below this line */ }
    return <Items quantity = {10} />
    { /* Change code above this line */ }
  }
};
```
