---
id: 5a24c314108439a4d4036189
title: Change Inline CSS Conditionally Based on Component State
challengeType: 6
forumTopicId: 301380
dashedName: change-inline-css-conditionally-based-on-component-state
---

# --description--

A questo punto, hai visto diversi esempi di rendering condizionale e l’uso di stili in linea. Ecco un altro esempio che combina entrambi questi argomenti. Puoi anche renderizzare il CSS condizionalmente in base allo stato di un componente React. Per farlo, controlli una condizione e, se questa è soddisfatta, modifichi l’oggetto degli stili assegnato agli elementi JSX nel metodo render.

Questo paradigma è importante da capire perché rappresenta un cambiamento radicale rispetto all’approccio più tradizionale di applicare gli stili modificando direttamente gli elementi del DOM (molto comune con jQuery, per esempio). In quell’approccio, devi tenere traccia di quando gli elementi cambiano e gestire direttamente la manipolazione effettiva. Può diventare difficile seguire i cambiamenti, rendendo potenzialmente imprevedibile l’interfaccia utente. Quando imposti un oggetto stile basato su una condizione, descrivi come l’interfaccia dovrebbe apparire in funzione dello stato dell’applicazione. C’è un flusso di informazioni chiaro che si muove in una sola direzione. Questo è il metodo preferito quando si scrivono applicazioni con React.

# --instructions--

L’editor del codice contiene un semplice componente di input controllato con un bordo stilizzato. Vuoi definire lo stile di questo bordo in rosso se l’utente digita più di 15 caratteri nel riquadro di input. Aggiungi una condizione per controllare questo e, se la condizione è valida, imposta lo stile del bordo dell’input su `3px solid red`. Puoi provarlo inserendo del testo nell’input.

# --hints--

Il componente `GateKeeper` dovrebbe renderizzare un elemento `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.find('div').length === 1;
  })()
);
```

Il componente `GateKeeper` dovrebbe essere inizializzato con una chiave di stato `input` impostata a una stringa vuota.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return mockedComponent.state().input === '';
  })()
);
```

Il componente `GateKeeper` dovrebbe renderizzare un tag `h3` e un tag `input`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('h3').length === 1 &&
      mockedComponent.find('input').length === 1
    );
  })()
);
```

Il tag `input` dovrebbe inizialmente avere uno stile `1px solid black` per la proprietà `border`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
    return (
      mockedComponent.find('input').props().style.border === '1px solid black'
    );
  })()
);
```

Il tag `input` dovrebbe essere stilizzato con un bordo `3px solid red` se il valore dell’input nello stato è più lungo di 15 caratteri.

```js
  const waitForIt = (fn) =>
    new Promise((resolve, reject) => setTimeout(() => resolve(fn()), 100));
  const mockedComponent = Enzyme.mount(React.createElement(GateKeeper));
  const simulateChange = (el, value) =>
    el.simulate('change', { target: { value } });
  let initialStyle = mockedComponent.find('input').props().style.border;
  const state_1 = () => {
    mockedComponent.setState({ input: 'this is 15 char' });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const state_2 = () => {
    mockedComponent.setState({
      input: 'A very long string longer than 15 characters.'
    });
    return waitForIt(() => mockedComponent.find('input').props().style.border);
  };
  const style_1 = await state_1();
  const style_2 = await state_2();
  assert(
    initialStyle === '1px solid black' &&
      style_1 === '1px solid black' &&
      style_2 === '3px solid red'
  );
```

# --seed--

## --seed-contents--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    // Change code below this line

    // Change code above this line
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```

# --solutions--

```jsx
class GateKeeper extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: ''
    };
    this.handleChange = this.handleChange.bind(this);
  }
  handleChange(event) {
    this.setState({ input: event.target.value })
  }
  render() {
    let inputStyle = {
      border: '1px solid black'
    };
    if (this.state.input.length > 15) {
      inputStyle.border = '3px solid red';
    };
    return (
      <div>
        <h3>Don't Type Too Much:</h3>
        <input
          type="text"
          style={inputStyle}
          value={this.state.input}
          onChange={this.handleChange} />
      </div>
    );
  }
};
```
