---
id: 5a24c314108439a4d403617c
title: Use the Lifecycle Method componentWillMount
challengeType: 6
forumTopicId: 301423
dashedName: use-the-lifecycle-method-componentwillmount
---

# --description--

I componenti React hanno diversi metodi speciali che offrono la possibilità di eseguire azioni in momenti specifici del ciclo di vita di un componente. Questi sono chiamati metodi del ciclo di vita, o hook del ciclo di vita, e ti permettono di intercettare i componenti in determinati momenti. Questo può avvenire prima che vengano renderizzati, prima che si aggiornino, prima che ricevano props, prima che vengano smontati, e così via. Ecco una lista di alcuni dei principali metodi del ciclo di vita: `componentWillMount()` `componentDidMount()` `shouldComponentUpdate()` `componentDidUpdate()` `componentWillUnmount()` Le prossime lezioni copriranno alcuni degli usi base di questi metodi del ciclo di vita.

**Nota:** Il metodo del ciclo di vita `componentWillMount` sarà deprecato in una futura versione 16.X e rimosso nella versione 17. Scopri di più in questo <a href="https://www.freecodecamp.org/news/how-to-safely-use-reacts-life-cycles-with-fiber-s-async-rendering-fd4469ebbd8f/" target="_blank" rel="noopener noreferrer nofollow">articolo</a>

# --instructions--

Il metodo `componentWillMount()` viene chiamato prima del metodo `render()` quando un componente viene montato nel DOM. Registra qualcosa nella console all'interno di `componentWillMount()` — potresti voler tenere aperta la console del browser per vedere l'output.

# --hints--

`MyComponent` dovrebbe renderizzare un elemento `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(MyComponent));
    return mockedComponent.find('div').length === 1;
  })()
);
```

`console.log` dovrebbe essere chiamato in `componentWillMount`.

```js
assert(
  (function () {
    const lifecycle = React.createElement(MyComponent)
      .type.prototype.componentWillMount.toString()
      .replace(/ /g, '');
    return lifecycle.includes('console.log(');
  })()
);
```

# --seed--

## --seed-contents--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line

    // Change code above this line
  }
  render() {
    return <div />
  }
};
```

# --solutions--

```jsx
class MyComponent extends React.Component {
  constructor(props) {
    super(props);
  }
  componentWillMount() {
    // Change code below this line
    console.log('Component is mounting...');
    // Change code above this line
  }
  render() {
    return <div />
  }
};
```
