---
id: 5a24c314108439a4d4036160
title: Define an HTML Class in JSX
challengeType: 6
forumTopicId: 301393
dashedName: define-an-html-class-in-jsx
---

# --description--

Ora che ti stai abituando a scrivere JSX, potresti chiederti in cosa differisce dall’HTML.

Finora, potrebbe sembrare che HTML e JSX siano esattamente uguali.

Una differenza chiave in JSX è che non puoi più usare la parola `class` per definire le classi HTML. Questo perché `class` è una parola riservata in JavaScript. Invece, JSX usa `className`.

Infatti, la convenzione di denominazione per tutti gli attributi HTML e i riferimenti agli eventi in JSX diventa camelCase. Per esempio, un evento click in JSX è `onClick`, invece di `onclick`. Allo stesso modo, `onchange` diventa `onChange`. Anche se questa è una differenza sottile, è importante tenerla a mente andando avanti.

# --instructions--

Applica una classe `myDiv` all’elemento `div` fornito nel codice JSX.

# --hints--

La costante `JSX` dovrebbe restituire un elemento `div`.

```js
assert.strictEqual(JSX.type, 'div');
```

L’elemento `div` dovrebbe avere una classe `myDiv`.

```js
assert.strictEqual(JSX.props.className, 'myDiv');
```

# --seed--

## --seed-contents--

```jsx
const JSX = (
  <div>
    <h1>Add a class to this div</h1>
  </div>
);
```

# --solutions--

```jsx
const JSX = (
<div className = 'myDiv'>
  <h1>Add a class to this div</h1>
</div>);
```
