---
id: 5a24c314108439a4d403616b
title: Use Default Props
challengeType: 6
forumTopicId: 301418
dashedName: use-default-props
---

# --description--

React offre anche la possibilità di impostare props predefinite. Puoi assegnare props predefinite a un componente come proprietà dello stesso componente e React assegna la prop predefinita se necessario. Questo ti permette di specificare quale dovrebbe essere il valore di una prop se non viene fornito esplicitamente un valore. Per esempio, se dichiari `MyComponent.defaultProps = { location: 'San Francisco' }`, hai definito una prop location impostata sulla stringa `San Francisco`, a meno che non specifichi diversamente. React assegna le props predefinite se le props sono undefined, ma se passi `null` come valore per una prop, questo rimarrà `null`.

# --instructions--

L'editor del codice mostra un componente `ShoppingCart`. Definisci le props predefinite su questo componente specificando una prop `items` con valore `0`.

# --hints--

Il componente `ShoppingCart` dovrebbe renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    return mockedComponent.find('ShoppingCart').length === 1;
  })()
);
```

Il componente `ShoppingCart` dovrebbe avere una prop predefinita `{ items: 0 }`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(ShoppingCart));
    mockedComponent.setProps({ items: undefined });
    return mockedComponent.find('ShoppingCart').props().items === 0;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};
// Change code below this line
```

# --solutions--

```jsx
const ShoppingCart = (props) => {
  return (
    <div>
      <h1>Shopping Cart Component</h1>
    </div>
  )
};

// Change code below this line
ShoppingCart.defaultProps = {
  items: 0
}
```
