---
id: 5a24c314108439a4d4036170
title: Crea un componente con stato
challengeType: 6
forumTopicId: 301391
dashedName: create-a-stateful-component
---

# --description--

Uno degli argomenti più importanti in React è `state`. Lo stato consiste in qualsiasi dato che la tua applicazione deve conoscere e che può cambiare nel tempo. Vuoi che le tue app rispondano ai cambiamenti di stato e mostrino un’interfaccia aggiornata quando necessario. React offre una soluzione efficace per la gestione dello stato nelle applicazioni web moderne.

Crei lo stato in un componente React dichiarando una proprietà `state` nella classe del componente nel suo `constructor`. Questo inizializza il componente con `state` quando viene creato. La proprietà `state` deve essere impostata su una `object` JavaScript. La dichiarazione è simile a questa:

```jsx
this.state = {

}
```

Hai accesso all’oggetto `state` per tutta la durata del componente. Puoi aggiornarlo, renderizzarlo nell’interfaccia e passarlo come props ai componenti figli. L’oggetto `state` può essere complesso o semplice quanto ti serve. Nota che devi creare un componente di classe estendendo `React.Component` per poter creare `state` in questo modo.

# --instructions--

Nel codice c’è un componente che sta cercando di renderizzare una proprietà `firstName` dal suo `state`. Tuttavia, non è definito alcun `state`. Inizializza il componente con `state` nel `constructor` e assegna il tuo nome a una proprietà di `firstName`.

# --hints--

`StatefulComponent` dovrebbe esistere e renderizzare.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return mockedComponent.find('StatefulComponent').length === 1;
  })()
);
```

`StatefulComponent` dovrebbe renderizzare un elemento `div` e un elemento `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h1').length === 1
    );
  })()
);
```

Lo stato di `StatefulComponent` dovrebbe essere inizializzato con una proprietà `firstName` impostata su una stringa.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return (
      typeof initialState === 'object' && typeof initialState.firstName === 'string'
    );
  })()
);
```

La proprietà `firstName` nello stato di `StatefulComponent` dovrebbe essere renderizzata nell’elemento `h1`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(
      React.createElement(StatefulComponent)
    );
    const initialState = mockedComponent.state();
    return mockedComponent.find('h1').text() === initialState.firstName;
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<StatefulComponent />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    // Only change code below this line

    // Only change code above this line
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```

# --solutions--

```jsx
class StatefulComponent extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      firstName: 'freeCodeCamp!'
    }
  }
  render() {
    return (
      <div>
        <h1>{this.state.firstName}</h1>
      </div>
    );
  }
};
```
