---
id: 5a24c314108439a4d403616e
title: Consultare le props usando this.props
challengeType: 6
forumTopicId: 301375
dashedName: access-props-using-this-props
---

# --description--

Le ultime sfide hanno trattato i modi base per passare le props ai componenti figli. Ma cosa succede se il componente figlio a cui stai passando una prop è un componente classe ES6, invece di un componente funzionale senza stato? Il componente classe ES6 usa una convenzione leggermente diversa per consultare le props.

Ogni volta che ti riferisci a un componente classe al suo interno, usi la parola chiave `this`. Per consultare le props all’interno di un componente classe, devi anteporre al codice che usi per consultarle `this`. Per esempio, se un componente classe ES6 ha una prop chiamata `data`, scrivi `{this.props.data}` in JSX.

# --instructions--

Renderizza un’istanza del componente `Welcome` nel componente genitore `App`. Qui, dai a `Welcome` una prop chiamata `name` e assegnale il valore di una stringa. All’interno del figlio, `Welcome`, consulta la prop `name` all’interno dei tag `strong`.

# --hints--

Il componente `App` dovrebbe restituire un singolo elemento `div`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.children().type() === 'div';
  })()
);
```

Il figlio di `App` dovrebbe essere il componente `Welcome`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.children().childAt(0).name() === 'Welcome'
    );
  })()
);
```

Il componente `Welcome` dovrebbe avere una prop chiamata `name`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return mockedComponent.find('Welcome').props().name;
  })()
);
```

Il componente `Welcome` dovrebbe mostrare la stringa che passi come prop `name` all’interno dei tag `strong`.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(App));
    return (
      mockedComponent.find('strong').text() ===
      mockedComponent.find('Welcome').props().name
    );
  })()
);
```

# --seed--

## --after-user-code--

```jsx
ReactDOM.render(<App />, document.getElementById('root'))
```

## --seed-contents--

```jsx
class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome />
            { /* Change code above this line */ }
        </div>
    );
  }
};

class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong></strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};
```

# --solutions--

```jsx
class Welcome extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
          { /* Change code below this line */ }
          <p>Hello, <strong>{this.props.name}</strong>!</p>
          { /* Change code above this line */ }
        </div>
    );
  }
};

class App extends React.Component {
  constructor(props) {
    super(props);

  }
  render() {
    return (
        <div>
            { /* Change code below this line */ }
            <Welcome name="Quincy"/>
            { /* Change code above this line */ }
        </div>
    );
  }
};
```
