---
id: 69a92f0b9cc04bb0d5327bb5
title: スタックを実装する
challengeType: 26
dashedName: implement-a-stack
---

# --description--

このラボでは、関数を使ってスタックのデータ構造を実装します。スタックは後入れ先出し（LIFO）のデータ構造で、要素は上から追加および削除されます。

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `initStack` 関数を用意し、空の配列がセットされた `collection` プロパティを持つオブジェクトを返してください。`collection` の最後の要素がスタックのトップを表します。
1. `push` 関数を用意し、スタックオブジェクトと要素を引数に取り、スタックの `collection` の末尾に要素を追加してスタックのトップに置いてください。
1. `pop` 関数を用意し、スタックオブジェクトを引数に取り、スタックのトップ要素を削除して返してください。要素がなければ `undefined` を返します。
1. `peek` 関数を用意し、スタックオブジェクトを引数に取り、スタックのトップ要素を削除せずに返してください。要素がなければ `undefined` を返します。
1. `isEmpty` 関数を用意し、スタックオブジェクトを引数に取り、スタックに要素がなければ `true`、そうでなければ `false` を返してください。
1. `clear` 関数を用意し、スタックオブジェクトを引数に取り、スタックの `collection` からすべての要素を削除してください。

**注意:** ほとんどのテストは `initStack` と `push` 関数に依存しています。まずそれらを実装してください。`pop`、`peek`、`isEmpty`、`clear` のテストはスタックに要素を追加する必要があります。

# --hints--

`push` 関数を用意してください。

```js
assert.isFunction(push);
```

`push` 関数は要素をスタックのトップに追加する必要があります。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

`pop` 関数を用意してください。

```js
assert.isFunction(pop);
```

`pop` 関数はスタックのトップ要素を削除して返す必要があります。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
assert.strictEqual(pop(stack), 'third');
assert.deepEqual(stack.collection, ['first', 'second']);
assert.strictEqual(pop(stack), 'second');
assert.deepEqual(stack.collection, ['first']);
```

スタックが空の場合、`pop` 関数は `undefined` を返す必要があります。

```js
const stack = initStack();
assert.isUndefined(pop(stack));
```

`pop` 関数は偽値を正しく返す必要があります。

```js
const stack = initStack();
push(stack, '');
assert.strictEqual(pop(stack), '');
assert.lengthOf(stack.collection, 0);
```

`peek` 関数を用意してください。

```js
assert.isFunction(peek);
```

`peek` 関数はスタックのトップ要素を削除せずに返す必要があります。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
assert.strictEqual(peek(stack), 'second');
assert.strictEqual(peek(stack), 'second');
assert.deepEqual(stack.collection, ['first', 'second']);
```

スタックが空の場合、`peek` 関数は `undefined` を返す必要があります。

```js
const stack = initStack();
assert.isUndefined(peek(stack));
```

`peek` 関数は偽値を正しく返す必要があります。

```js
const stack = initStack();
push(stack, 0);
assert.strictEqual(peek(stack), 0);
```

`isEmpty` 関数を用意してください。

```js
assert.isFunction(isEmpty);
```

`isEmpty` 関数は空のスタックに対して `true` を返す必要があります。

```js
const stack = initStack();
assert.isTrue(isEmpty(stack));
```

`isEmpty` 関数は空でないスタックに対して `false` を返す必要があります。

```js
const stack = initStack();
push(stack, 'element');
assert.isFalse(isEmpty(stack));
```

トップ要素が偽値の場合でも、`isEmpty` 関数は `false` を返す必要があります。

```js
const stack = initStack();
push(stack, false);
assert.isFalse(isEmpty(stack));
```

`clear` 関数を持っている必要があります。

```js
assert.isFunction(clear);
```

`clear` 関数はスタックからすべての要素を削除する必要があります。

```js
const stack = initStack();
push(stack, 'first');
push(stack, 'second');
push(stack, 'third');
clear(stack);
assert.lengthOf(stack.collection, 0);
```

# --seed--

## --seed-contents--

```js
```

# --solutions--

```js
function initStack() {
  return {
    collection: []
  };
}

function print(stack) {
  console.log(stack.collection);
}

function push(stack, element) {
  stack.collection.push(element);
}

function pop(stack) {
  return stack.collection.pop();
}

function peek(stack) {
  return stack.collection[stack.collection.length - 1];
}

function isEmpty(stack) {
  return stack.collection.length === 0;
}

function clear(stack) {
  stack.collection = [];
}
```
