---
id: 66eaddd04a9e533fba689001
title: 设计一个博客文章卡片
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

在这个实验中，你将通过创建一个博客帖子卡片来练习如何设置背景和边框样式。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应有一个 `div` 元素，类名为 `blog-post-card`，用于容纳所有卡片元素。
2. 在 `.blog-post-card` 元素内，你应该一个图像元素，该元素具备有效的 `alt` 属性及文本，类名为 `post-img`。 你可以使用 `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` 作为图像的 `src` 属性的值。
3. 你应该在 `.blog-post-card` 元素内有一个带有类 `post-content` 的 `div`，其中包含以下内容：
    - 一个 `h2` 元素，类名为 `post-title`，包含你的博客文章的标题文本。
    - 一个 `p` 元素，类名为 `post-excerpt`，包含概述你的博客文章的文本。
    - 一个 `a` 元素， 类名为 `read-more`，包含文本 `Read More`。
4. 你应该将 `body` 元素的背景颜色设置为非白色的值。
5. 你应该将以下样式应用于 `.blog-post-card` 元素：
    - 白色背景。
    - 圆角。
    - 自定义的宽度。
    - 自定义的文本对齐方式。
6. 设置`.post-img` 元素的样式，以便图像能够填满卡片的整个宽度并有一个底部边框。
7. 设置`.post-content` 元素的样式，以便在卡片内部有内边距。
8. `.post-title` 和 `.post-excerpt` 元素的文本颜色应该不是默认的，所有边都有外边距。
9. `.read-more` 元素的样式应该被设置为像一个按钮，包含：
    - 非默认的文本颜色。
    - 背景颜色。
    - 所有边的外边距。
    - 所有边的内边距。
    - 圆角。
    - `display` 属性设置为 `inline-block`。
    - 改变背景颜色的悬停效果。

**注意：** 确保在 HTML 中链接样式表以应用 CSS。

# --hints--

你应该有一个带有 `blog-post-card` 类的 `div` 元素。

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

你应该有一个 `class` 为 `post-img` 的 `img` 元素。 确保你的图像有一个带有文本的 `alt` 属性和一个带有值的 `src` 属性。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

你应该有一个带有 `post-content` 类的 `div` 元素。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

你应该有一个带有 `post-title` 类的 `h2`元素。 确保博客文章标题有一些文字内容。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

你应该有一个带有 `post-excerpt` 类的 `p` 元素。 确保它有一些文字内容来概括你的博客文章。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

你应该有一个 `class` 是 `read-more` 的 `a` 元素。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

你的 `.read-more` 元素应该包含文本 `Read More`。

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

你的 `.blog-post-card` 元素应该有一个 有值的 `border-radius` 属性（不应该是 `0` 或负值）。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

你的 `.blog-post-card` 元素应该有白色的背景。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

你的 `body` 元素应具有非白色的背景色。

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

你应该以 `.blog-post-card` 为目标并设置它的 `width` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

你应该以 `.blog-post-card` 为目标并设置它的 `text-align` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

你应该以 `.post-content` 为目标并设置它的 `padding` 属性。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

你的 `.read-more` 元素应该具有悬停效果。

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

你应该以 `.read-more` 为目标并设置它的 `color` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

你应该以 `.read-more` 为目标并设置它的 `background-color` 属性。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

你应该以 `.read-more` 为目标并设置它的 `margin` 属性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

你应该以 `.read-more` 为目标并设置它的 `display` 属性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

你应该以 `.read-more` 为目标并设置它的 `border-radius` 属性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

你应该以 `.read-more` 为目标并设置它的 `padding` 属性。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

你的 `.post-img` 元素应该填充卡片的整个 `width`。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

你的 `.post-img` 元素应该有一个 `border-bottom` 值。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

你的 `.post-title` 和 `.post-excerpt` 元素应该有外边距和非默认的文本颜色。

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
