---
id: 695774002591bbc5f8cf3e53
title: بناء دالة تطبيق الخصم
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

في هذا المختبر ستكتب دالة تحسب السعر النهائي لعنصر بعد تطبيق نسبة خصم.

على سبيل المثال، إذا كان سعر العنصر `50` وتم تطبيق خصم `20`، فإن مبلغ الخصم هو `10`، والسعر النهائي هو `40`.

**الهدف:** استكمل قصص المستخدم أدناه لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `apply_discount`.

2. يجب أن تأخذ دالة `apply_discount` معلمتين بالضبط: `price` و `discount`.

3. إذا لم يكن `price` رقمًا (`int` أو `float`)، يجب أن تُرجع الدالة السلسلة النصية `The price should be a number`.

4. إذا لم يكن `discount` رقمًا (`int` أو `float`)، يجب أن تُرجع الدالة السلسلة النصية `The discount should be a number`.

5. إذا كان `price` أقل من أو يساوي `0`، يجب أن تُرجع الدالة السلسلة النصية `The price should be greater than 0`.

6. إذا كان `discount` أقل من `0` أو أكبر من `100`، يجب أن تُرجع الدالة السلسلة النصية `The discount should be between 0 and 100`.

7. إذا كان كلا المُدخلين صالحين، يجب أن تحسب الدالة الخصم كنسبة مئوية من السعر.

8. يجب أن تُرجع الدالة السعر النهائي بعد تطبيق الخصم.

# --hints--

يجب أن يكون لديك دالة باسم `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

يجب أن تأخذ دالة `apply_discount` معلمتين: `price` و `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

عند استدعاء `apply_discount` مع `price` (المعلمة الأولى) التي ليست رقمًا (`int` أو `float`)، يجب أن تُرجع `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

عند استدعاء `apply_discount` مع `discount` (المعلمة الثانية) التي ليست رقمًا (`int` أو `float`)، يجب أن تُرجع `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

عند استدعاء `apply_discount` مع `price` أقل من أو يساوي `0`، يجب أن تُرجع `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

عند استدعاء `apply_discount` مع `discount` أقل من `0` أو أكبر من `100`، يجب أن تُرجع `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` يجب أن تُرجع `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` يجب أن تُرجع `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` يجب أن تُرجع `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

عند استدعاء `apply_discount` مع خصم `100`، يجب أن تُرجع `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` يجب أن تُرجع `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
