---
id: 690ddb5e6ee94eb73ed172b0
title: Depura un Formulario de Donación
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Una organización benéfica local ha creado un sitio web con un formulario de donación, pero hay varios problemas que necesitan ser corregidos. El formulario no es accesible y tiene algunos errores de sintaxis en HTML.

Tu trabajo es corregir todos los errores para que el formulario funcione correctamente y sea accesible para todos los usuarios. Completa los elementos en las historias de usuario a continuación y haz clic en "Run the Tests" para ver si corregiste todos los errores.

**Historias de usuario:**

1. Los elementos `input` son elementos vacíos y no deben tener etiquetas de cierre. Elimina todas las etiquetas de cierre `</input>` del formulario.
2. Agrega elementos `label` para cada campo de entrada del formulario para que los usuarios sepan para qué sirve cada campo. El texto del label debe coincidir con lo que actualmente está al lado de cada entrada.
3. El tipo de entrada `Email Address:` debería ser un `email` en lugar de `text`.
4. Debes asociar cada elemento `label` con su elemento `input` correspondiente usando el atributo `for` en el `label` y un atributo `id` coincidente en el `input`.
5. Agrega el atributo `required` a los campos de entrada de texto, correo electrónico y número (pero no al checkbox ni al botón de envío) para asegurar que los usuarios completen la información requerida.

# --hints--

No debes tener ninguna etiqueta de cierre `</input>` en tu código.

```js
assert.notMatch(code, /<\/input>/);
```

Debes tener exactamente cinco elementos `input` en tu formulario.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Debes tener exactamente cuatro elementos `label` en tu formulario.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Tu primer `label` debe tener el texto `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Tu primer `input` debe tener un atributo `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Tu segundo `label` debe tener el texto `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Tu segundo `input` debe tener un atributo `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Tu tercer `label` debe tener el texto `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Tu tercer `input` debe tener un atributo `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Tu cuarto `label` debe tener el texto `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Debes asociar cada elemento `input` con un elemento `label` usando los atributos `for` e `id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Tu `input` para la dirección de correo electrónico debe ser del tipo `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Tu `input` de casilla de verificación no debe tener un atributo `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Tu `button` de envío no debe tener un atributo `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
