---
id: 699f513e01eb1768cc5df726
title: Passo 17
challengeType: 1
dashedName: step-17
---

# --description--

Agora que o parâmetro `exponent` está marcado como opcional, você pode chamar a função com apenas um argumento.

Atualize sua chamada da função `raiseTo` para passar somente o número `3`.

# --hints--

Sua função `raiseTo` deve ser chamada com apenas um argumento `3`.

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
