---
id: 699f4861dd5b5f164e1bda4a
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Outro benefício de usar TypeScript é garantir que os tipos de retorno sejam aplicados e correspondam ao que você espera. Às vezes, o TypeScript pode inferir o tipo de retorno de uma função, mas você também pode especificá-lo explicitamente.

Para ilustrar isso, altere o tipo de retorno da função `getAverage` para a string `"something"`. 

# --hints--

Sua função `getAverage` deve retornar uma string.

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

Sua função `getAverage` deve retornar a string `"something"`.

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
