---
id: 699f3738921fb79a651ded56
title: passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora é hora de chamar sua função.

Crie uma variável chamada `result` e atribua a ela a chamada da função `square` com um argumento `5`.

Abaixo disso, registre o valor de `result` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável chamada `result`.

```js
assert.isDefined(result);
```

Sua variável `result` deve ser um número.

```js
assert.isNumber(result);
```

Sua variável `result` deve ser o quadrado de `5`.

```js
assert.strictEqual(result, 25);
```

Você deve registrar a variável `result` no console.

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
