---
id: 699f3f636aa840b7201aa37c
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Se você abrir o console, deverá ver a seguinte mensagem de erro do TypeScript:

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

Agora que o parâmetro está explicitamente tipado como número, o TypeScript consegue detectar o erro de tentar passar um argumento string para a função `square`.

Altere o argumento passado para a função `square` para o número `5` para corrigir o erro e ver a saída correta no console.

# --hints--

Sua chamada da função `square` deve ter um argumento `5` em vez de `"something"`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
