---
id: 699f4ec939f10da62fb56c3e
title: Passo 15
challengeType: 1
dashedName: step-15
---

# --description--

Agora é hora de testar sua função.

Crie uma variável chamada `raisedResult` e atribua a ela o resultado da chamada de `raiseTo` com os argumentos `2` e `3`.

Depois, registre `raisedResult` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável chamada `raisedResult`.

```js
assert.isDefined(raisedResult);
```

Sua variável `raisedResult` deve ser um número.

```js
assert.isNumber(raisedResult);
```

Você deve registrar `raisedResult` no console.

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
