---
id: 699f1d72e906e4616155290d
title: passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Nesta oficina, você vai continuar aprendendo sobre anotações de tipo e como elas funcionam com funções construindo uma caixa de ferramentas matemática com segurança de tipos.

Comece criando uma função chamada `square` que aceita um parâmetro e retorna o quadrado de um número.

# --hints--

Você deve ter uma função chamada `square`.

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

Sua função `square` deve ter um parâmetro.

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

Sua função `square` deve retornar o quadrado de um número.

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
