---
id: 699f49669ea51a31d2ba89d8
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Como você lembra das lições anteriores, você pode tipar explicitamente o tipo de retorno de uma função assim:

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

Neste exemplo, o tipo de retorno da função `circleArea` é `number`. Se você tentar retornar um tipo diferente, o TypeScript exibirá um erro.

Atualize sua função `getAverage` para tipar explicitamente o tipo de retorno como `number`. 

# --hints--

Sua função `getAverage` deve ter um tipo de retorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
