---
id: 699f4a3e7b4c6bfa538cc4f7
title: Passo 13
challengeType: 1
dashedName: step-13
---

# --description--

Se você abrir o console, verá a seguinte mensagem de erro:

```md
Type 'string' is not assignable to type 'number'.
```

Para resolver esse erro, você precisa mudar o retorno para `sum / numbers.length`. 

# --hints--

Sua função `getAverage` deve retornar um número.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Sua função `getAverage` deve retornar a média dos números no array.

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
