---
id: 699f434e294a9e11cf4f65f0
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

Agora é hora de adicionar mais um auxiliar matemático ao seu projeto toolkit.

Comece criando uma função chamada `getAverage` que recebe um parâmetro. Sua função deve retornar a média de um array de números.

Nos próximos passos, você vai adicionar as anotações de tipo para essa função.

# --hints--

Você deve ter uma função chamada `getAverage`.

```js
assert.isFunction(getAverage);
```

Sua função deve receber um parâmetro.

```js
assert.strictEqual(getAverage.length, 1);
```

Sua função deve retornar um número.

```js
assert.isNumber(getAverage([1, 2, 3]));
```

Sua função deve retornar o resultado correto para qualquer array de números.

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
