---
id: 699f52220de88c8091032076
title: Passo 18
challengeType: 1
dashedName: step-18
---

# --description--

Se você abrir o console, o resultado será `NaN` porque o segundo argumento para a `exponent` nunca foi fornecido. Existem algumas atualizações necessárias para essa função `raiseTo` que vão corrigir isso.

A primeira mudança é retornar `base ** 2`.

No próximo passo, você verá como lidar com o parâmetro `exponent` opcional.

# --hints--

Sua função `raiseTo` deve retornar um número.

```js
assert.isNumber(raiseTo(4));
```

Sua função `raiseTo` deve retornar o valor de `base` elevado à potência de `2`.

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
