---
id: 699f4b903bf5427a599f791d
title: Passo 14
challengeType: 1
dashedName: step-14
---

# --description--

Nos Ășltimos passos, vocĂȘ vai criar mais uma funĂ§ĂŁo para continuar praticando anotaĂ§Ă”es de tipo com funĂ§Ă”es.

Comece criando uma funĂ§ĂŁo chamada `raiseTo` que recebe dois parĂąmetros chamados `base` e `exponent`. Ambos os parĂąmetros devem ser do tipo `number`. A funĂ§ĂŁo deve retornar o resultado de elevar `base` Ă  potĂȘncia de `exponent`. O tipo de retorno da funĂ§ĂŁo tambĂ©m deve ser anotado como `number`.

# --hints--

VocĂȘ deve ter uma funĂ§ĂŁo `raiseTo`.

```js
assert.isFunction(raiseTo);
```

Sua funĂ§ĂŁo `raiseTo` deve ter dois parĂąmetros, `base` e `exponent`.

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

Seus parĂąmetros `base` e `exponent` devem ser ambos do tipo `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

Sua funĂ§ĂŁo `raiseTo` deve retornar um nĂșmero.

```js
assert.isNumber(raiseTo(2, 3));
```

Sua funĂ§ĂŁo `raiseTo` deve retornar o resultado de elevar `base` Ă  potĂȘncia de `exponent`.

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

Sua funĂ§ĂŁo `raiseTo` deve ter uma anotaĂ§ĂŁo de tipo de retorno `number`.

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
