---
id: 699f52fac083214b6bf51de6
title: Passo 19
challengeType: 1
dashedName: step-19
---

# --description--

Neste último passo, você precisará adicionar uma verificação condicional para o parâmetro `exponent` na função `raiseTo`.

Se o argumento `exponent` for fornecido, a função deve retornar `base` elevado à potência de `exponent`. Se o argumento `exponent` não for fornecido, deve retornar `base ** 2`.

Com esse último passo, seu kit de ferramentas matemáticas com segurança de tipos estará completo!

# --hints--

Sua função `raiseTo` deve retornar `base ** exponent` quando o argumento `exponent` for fornecido.

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

Sua função `raiseTo` deve retornar `base ** 2` quando o argumento `exponent` não for fornecido.

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
