---
id: 699f412114c87221f16a99e3
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Outro benefício de usar TypeScript é que ele detecta erros se você passar argumentos demais ou de menos em uma chamada de função.

Para ilustrar isso, atualize sua chamada da função `square` para receber dois argumentos `5` e `10`.

# --hints--

Sua chamada da função `square` deve ter dois argumentos `5` e `10`.

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
