---
id: 699f41b8da20f816b5c7dece
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Se você abrir o console, deverá ver a seguinte mensagem de erro:

```md
Expected 1 arguments, but got 2.
```

Para resolver o erro, atualize sua chamada da função `square` para aceitar apenas um argumento do tipo `5`.

# --hints--

Sua chamada da função `square` deve ter apenas um argumento do tipo `5`.

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
