---
id: 699f458e47fb88c756fdba81
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Neste momento, o parâmetro `numbers` tem um tipo `any` implícito. Para corrigir isso, você pode tipar explicitamente o parâmetro como um array de números.

Lembre-se das lições anteriores, que você pode tipar um array de números assim:

```ts
const numbers: number[] = [1, 2, 3];
```

Atualize sua função `getAverage` para tipar explicitamente o parâmetro `numbers` como um array de números.

# --hints--

Sua função `getAverage` deve ter um parâmetro chamado `numbers` que seja explicitamente tipado como um array de números.

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
