---
id: 699f3cb8d057f29954205d7d
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Como você lembra das lições anteriores, é possível adicionar anotações de tipo aos parâmetros de função assim:

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

Neste exemplo, o TypeScript exibirá um erro se você tentar fazer uma chamada da função `add` com argumentos que não sejam números.

Atualize sua função `square` para incluir uma anotação de tipo para o parâmetro `num`. O tipo deve ser `number`.

# --hints--

Você deve adicionar uma anotação de tipo ao parâmetro `num` na função `square`. O tipo deve ser um `number`.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
