---
id: 699f4639ae0016caaf81b7c6
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Agora é hora de testar sua função.

Comece criando uma variável chamada `avgResult` e atribuindo a ela a chamada da função `getAverage`. O argumento para essa chamada de função deve ser `[2, 14, 26, 8]`.

Depois, registre o valor de `avgResult` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

Você deve ter uma variável chamada `avgResult`.

```js
assert.isDefined(avgResult);
```

Você deve atribuir `getAverage([2, 14, 26, 8])` à sua variável `avgResult`.

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

Sua variável `avgResult` deve ser um número.

```js
assert.isNumber(avgResult);
```

Você deve registrar `avgResult` no console.

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
