---
id: 699f38937b053e4679878ff1
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Ao trabalhar com JavaScript puro, nada impede que você faça uma chamada da função `square` com uma string ou um array, o que levaria a resultados inesperados.

Para ilustrar isso, altere sua chamada de função para usar a string `"something"` em vez do número `5`.

Na próxima lição, você vai aprender a usar TypeScript para adicionar segurança de tipos ao seu código e evitar esse tipo de erro.

# --hints--

Sua chamada da função `square` deve ter um argumento do tipo `"something"`.

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
