---
id: 699f5063d2393c674b1f45dc
title: Passo 16
challengeType: 1
dashedName: step-16
---

# --description--

Como você se lembra de um passo anterior, se você chamar uma função com o número errado de argumentos, o TypeScript gerará um erro.

Existe uma forma de marcar parâmetros como opcionais adicionando um `?` após o nome do parâmetro. Isso permite chamar a função com menos argumentos do que parâmetros.

Aqui está um exemplo:

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

Neste exemplo, se o parâmetro `greeting` não for fornecido ao chamar a função `greet`, ele será `undefined`, e a função retornará uma saudação padrão.

Atualize sua função `raiseTo` para tornar o parâmetro `exponent` opcional.

# --hints--

Sua função `raiseTo` deve ter o parâmetro `exponent` marcado como opcional.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
