---
id: 686b9720ee1d032bd77a480a
title: Depura un Validador de ISBN
challengeType: 27
dashedName: debug-an-isbn-validator
---

# --description--

El ISBN (International Standard Book Number) es un identificador único asignado a libros comerciales. Puede tener 10 o 13 dígitos, y el último dígito es un dígito de control calculado a partir de los otros dígitos.

Camperbot ha intentado construir su propia estructura validadora de ISBN. Sin embargo, ha cometido algunos errores en el proceso.

En este laboratorio, corregirás el código existente y harás que funcione correctamente.

Comportamiento esperado:  
Cuando el usuario ejecute el programa, mostrará el mensaje `Enter ISBN and length: `. El usuario puede ingresar el código ISBN que desea validar en el formato `ISBN,length`. El código ISBN no debe contener guiones, seguido de su longitud (`10` o `13`), separados por una coma.

Ejemplos de entradas:  
`1530051126,10` para códigos ISBN-10.  
`9781530051120,13` para códigos ISBN-13.

Cómo encontrar el dígito de verificación:

- No tienes que conocer la lógica detallada de cálculo en este laboratorio. Las funciones `calculate_check_digit_10` y `calculate_check_digit_13` se encargarán del cálculo y devolverán el dígito de control esperado en forma de cadena.
- El dígito de verificación para códigos ISBN-10 puede ser un número de `0` a `9` o una letra mayúscula `X`.
- El dígito de verificación para códigos ISBN-13 puede ser un número de `0` a `9`.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes corregir el `IndentationError` en el código actual.
1. Incluso si el usuario no ingresa un valor separado por comas, el programa debe manejar el `IndexError` sin fallar.
1. Cuando el usuario no ingresa un valor separado por comas, debería ver el mensaje `Enter comma-separated values.` en la consola, y el programa debería terminar.
1. Incluso si el usuario ingresa un valor no numérico para la longitud, el programa debe manejar el `ValueError` sin fallar.
1. Cuando el usuario ingresa un valor no numérico para la longitud, debería ver el mensaje `Length must be a number.` en la consola, y el programa debería terminar.
1. Debes corregir el error de desbordamiento en la función `validate_isbn`.
1. Debes corregir el `TypeError` en el código actual que ocurre cuando el usuario ingresa un código ISBN válido.
1. Debes corregir el `IndexError` en el código actual cuando el usuario ingresa un código ISBN válido.
1. Incluso si el usuario ingresa un código ISBN incorrecto con caracteres que no sean números, el programa debe manejar el `ValueError` que ocurre sin fallar.
1. Cuando el usuario ingresa un código ISBN incorrecto con caracteres que no sean números, debería ver el mensaje `Invalid character was found.` en la consola.
1. Cuando el usuario ingresa `1530051126,10`, debería ver el mensaje `Valid ISBN Code.`
1. Cuando el usuario ingresa `9781530051120,13`, debería ver el mensaje `Valid ISBN Code.`

**Importante**: necesitarás comentar la llamada a `main()` en el espacio global para que las pruebas se ejecuten correctamente.

Cuando completes el proyecto, el usuario debería ver los siguientes mensajes dependiendo de los valores que ingrese.

| Código ISBN | Longitud | Mensaje | Ejemplo de entrada |
|-----------|--------|---------| ------------- |
| Válido | Válida | `Valid ISBN Code.` | `1530051126,10` |
| Número inválido | Válida | `Invalid ISBN Code.` | `1530051125,10` |
| Sin longitud especificada o blanco | Válida | `ISBN-10 code should be 10 digits long.` o `ISBN-13 code should be 13 digits long.`, dependiendo de la longitud de entrada. | `9781530051120,10` o `1530051126,13` |
| Contiene caracteres no numéricos (excepto por el dígito de verificación) | Válida | `Invalid character was found.` | `15-0051126,10` |
| Cualquiera | Número inválido | `Length should be 10 or 13.` | `1530051126,9` |
| Cualquiera | Contiene caracteres no númericos o blanco | `Length must be a number.` | `1530051125,A` |
| Sin separación por la coma | Sin separación por la coma | `Enter comma-separated values.` | `1530051125` |

Puedes usar los siguientes valores para probar tu código manualmente si quieres.

Ejemplos de entradas para códigos ISBN-10 válidos:  
`1530051126,10`  
`9971502100,10`  
`080442957X,10`

Ejemplos de entradas para códigos ISBN-13 válidos:  
`9781530051120,13`  
`9781947172104,13`

# --hints--

Debes comentar la llamada a la función `main` para permitir que el resto de las pruebas funcionen correctamente.

```js
({
  test: () => runPython('assert not _Node(_code).has_call("main()")')
})
```

Debes tener una función `validate_isbn`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("validate_isbn")')
})
```

Debes tener una función `calculate_check_digit_10`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_10")')
})
```

Debes tener una función `calculate_check_digit_13`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("calculate_check_digit_13")')
})
```

Debes tener una función `main`.

```js
({
  test: () => runPython('assert _Node(_code).has_function("main")')
})
```

Cuando el usuario ingresa un valor que no es un valor separado por comas, deberías ver el mensaje `Enter comma-separated values.` en la consola.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Cuando el usuario ingresa un valor no numérico para la longitud, deberías ver el mensaje `Length must be a number.` en la consola.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,a"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print
try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Cuando el usuario ingresa un código ISBN incorrecto con caracteres que no sean números, deberías ver el mensaje `Invalid character was found.` en la consola.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "978IS3OOSII2O,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Cuando el usuario ingresa `1530051126,10`, deberías ver el mensaje `Valid ISBN Code.` en la consola.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Cuando el usuario ingresa `9781530051120,13`, deberías ver el mensaje `Valid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Cuando el usuario ingresa `1530051125,10`, deberías ver el mensaje `Invalid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid ISBN Code." in out

`)
}})
```

Cuando el usuario ingresa `9781530051120,10`, deberías ver el mensaje `ISBN-10 code should be 10 digits long.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781530051120,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-10 code should be 10 digits long." in out

`)
}})
```

Cuando el usuario ingresa `1530051126,13`, deberías ver el mensaje `ISBN-13 code should be 13 digits long.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "ISBN-13 code should be 13 digits long." in out

`)
}})
```

Cuando el usuario ingresa `15-0051126,10`, deberías ver el mensaje `Invalid character was found.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "15-0051126,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Invalid character was found." in out

`)
}})
```

Cuando el usuario ingresa `1530051126,9`, deberías ver el mensaje `Length should be 10 or 13.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051126,9"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length should be 10 or 13." in out

`)
}})
```

Cuando el usuario ingresa `1530051125,A`, deberías ver el mensaje `Length must be a number.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125,A"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Length must be a number." in out

`)
}})
```

Cuando el usuario ingresa `1530051125`, deberías ver el mensaje `Enter comma-separated values.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "1530051125"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Enter comma-separated values." in out

`)
}})
```

Cuando el usuario ingresa `9971502100,10`, deberías ver el mensaje `Valid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9971502100,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Cuando el usuario ingresa `080442957X,10`, deberías ver el mensaje `Valid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "080442957X,10"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

Cuando el usuario ingresa `9781947172104,13`, deberías ver el mensaje `Valid ISBN Code.`.

```js
({test: () => {runPython(`
built_in_print = print
out = []
def input(*args):
    return "9781947172104,13"
def custom_print(*args, **kwargs):
    call_args = [arg for arg in args]
    out.extend(call_args)
    
print = custom_print

try:
  main()
except:
  pass
finally:
  print = built_in_print

assert "Valid ISBN Code." in out

`)
}})
```

# --seed--

## --seed-contents--

```py
def validate_isbn(isbn, length):
    if len(isbn, length) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return
    main_digits = isbn[0:length]
    given_check_digit = isbn[length]
    main_digits_list = [int(digit) for digit in main_digits]
    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)
    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')
def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11
    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.
    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10
    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit
def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]
    length = int(values[1])
    if length == 10 or length == 13:
    validate_isbn(isbn, length)
    else:
    print('Length should be 10 or 13.')

main()
```

# --solutions--

```py
def validate_isbn(isbn, length):
    if len(isbn) != length:
        print(f'ISBN-{length} code should be {length} digits long.')
        return

    check_digit_index = length - 1

    main_digits = isbn[0:check_digit_index]
    given_check_digit = isbn[check_digit_index]

    try:
        main_digits_list = [int(digit) for digit in main_digits]
    except ValueError as e:
        print('Invalid character was found.')
        return

    # Calculate the check digit from other digits
    if length == 10:
        expected_check_digit = calculate_check_digit_10(main_digits_list)
    else:
        expected_check_digit = calculate_check_digit_13(main_digits_list)

    # Check if the given check digit matches with the calculated check digit
    if given_check_digit == expected_check_digit:
        print('Valid ISBN Code.')
    else:
        print('Invalid ISBN Code.')


def calculate_check_digit_10(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 9 digits by its corresponding weight (10 to 2) and sum up the results
    for index, digit in enumerate(main_digits_list):
        digits_sum += digit * (10 - index)
    # Find the remainder of dividing the sum by 11, then subtract it from 11
    result = 11 - digits_sum % 11

    # The calculation result can range from 1 to 11.
    # If the result is 11, use 0.
    # If the result is 10, use upper case X.
    # Use the value as it is for other numbers.
    if result == 11:
        expected_check_digit = '0'
    elif result == 10:
        expected_check_digit = 'X'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def calculate_check_digit_13(main_digits_list):
    # Note: You don't have to fully understand the logic in this function.

    digits_sum = 0
    # Multiply each of the first 12 digits by 1 and 3 alternately (starting with 1), and sum up the results
    for index, digit in enumerate(main_digits_list):
        if index % 2 == 0:
            digits_sum += digit * 1
        else:
            digits_sum += digit * 3
    # Find the remainder of dividing the sum by 10, then subtract it from 10
    result = 10 - digits_sum % 10

    # The calculation result can range from 1 to 10.
    # If the result is 10, use 0.
    # Use the value as it is for other numbers.
    if result == 10:
        expected_check_digit = '0'
    else:
        expected_check_digit = str(result)
    return expected_check_digit


def main():
    user_input = input('Enter ISBN and length: ')
    values = user_input.split(',')
    isbn = values[0]

    try:
        length = int(values[1])
    except ValueError as e:
        print('Length must be a number.')
        return
    except IndexError as e:
        print('Enter comma-separated values.')
        return

    if length == 10 or length == 13:
        validate_isbn(isbn, length)
    else:
        print('Length should be 10 or 13.')


# main()
```
