---
id: 67116c175a77cb64c5c10d49
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في الدروس السابقة، تعلمت كيف تعمل مع الدوال ذات الرتبة العليا مثل `map` و`filter` و`reduce`.

في هذه الورشة، ستبني تطبيق مدير مكتبة يمنحك فرصة لممارسة العمل مع دوال ذات رتب عليا مختلفة.

للبدء، أنشئ متغيرًا اسمه `library` وقم بتعيين مصفوفة فارغة له.

# --hints--

يجب أن تنشئ مصفوفة `library`.

```js
assert.isArray(library)
```

يجب أن تكون مصفوفة `library` فارغة.

```js
assert.isEmpty(library)
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
