---
id: 67116d7584d0b469b14579bf
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

داخل مصفوفة `library`، أنشئ كائنًا بالخصائص والقيم التالية:

| الخاصية | القيمة   |
| ----------- | ------- |
| `title` | `"Your Next Five Moves: Master the Art of Business Strategy"`|
| `author` | `"Patrick Bet-David and Greg Dinkin"`|
| `about` | `"A book on how to plan ahead"`|
| `pages` | `320` |

# --hints--

يجب أن يكون لديك كائن داخل مصفوفة `library`.

```js
assert.isObject(library[0])
```

يجب أن يحتوي كائنك على خاصية `title`.

```js
assert.property(library[0], "title")
```

يجب أن تكون قيمة خاصية `title` هي `"Your Next Five Moves: Master the Art of Business Strategy"`.

```js
assert.propertyVal(library[0], "title", "Your Next Five Moves: Master the Art of Business Strategy");
```

يجب أن يحتوي كائنك على خاصية `author`.

```js
assert.property(library[0], "author")
```

يجب أن تكون قيمة خاصية `author` هي `"Patrick Bet-David and Greg Dinkin"`.

```js
assert.propertyVal(library[0], "author", "Patrick Bet-David and Greg Dinkin");
```

يجب أن يحتوي كائنك على خاصية `about`.

```js
assert.property(library[0], "about")
```

يجب أن تكون قيمة خاصية `about` هي `"A book on how to plan ahead"`.

```js
assert.propertyVal(library[0], "about", "A book on how to plan ahead");
```

يجب أن يحتوي كائنك على خاصية `pages`.

```js
assert.property(library[0], "pages")
```

يجب أن تكون قيمة خاصية `pages` رقمًا.

```js
assert.isNumber(library[0]?.pages)
```

يجب تعيين خاصية `pages` إلى القيمة `320`.

```js
assert.propertyVal(library[0], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
--fcc-editable-region--
  
--fcc-editable-region--
];
```
