---
id: 67116d7584d0b469b14579c0
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

أنشئ كائنًا آخر داخل مصفوفة `library` بالخصائص والقيم التالية:

| الخاصية | القيمة   |
| ----------- | ------- |
| `title` | `"Atomic Habits"`|
| `author` | `"James Clear"`|
| `about` | `"A practical book about discarding bad habits and building good ones"`|
| `pages` | `320` |

# --hints--


يجب أن يكون لديك كائن كالعنصر الثاني في مصفوفة `library`.

```js
assert.isObject(library[1]);
```

يجب أن يحتوي كائنك على خاصية `title`.

```js
assert.property(library[1], "title")
```

يجب أن تكون قيمة خاصية `title` هي `"Atomic Habits"`.

```js
assert.propertyVal(library[1], "title", "Atomic Habits");
```

يجب أن يحتوي كائنك على خاصية `author`.

```js
assert.property(library[1], "author")
```

يجب أن تكون قيمة خاصية `author` هي `"James Clear"`.

```js
assert.propertyVal(library[1], "author", "James Clear");
```

يجب أن يحتوي كائنك على خاصية `about`.

```js
assert.property(library[1], "about")
```

يجب أن تكون قيمة خاصية `about` هي `"A practical book about discarding bad habits and building good ones"`.

```js
assert.propertyVal(library[1], "about", "A practical book about discarding bad habits and building good ones");
```

يجب أن يحتوي كائنك على خاصية `pages`.

```js
assert.property(library[1], "pages")
```

يجب أن تكون قيمة خاصية `pages` رقمًا.

```js
assert.isNumber(library[1]?.pages)
```

يجب تعيين خاصية `pages` إلى `320`.

```js
assert.propertyVal(library[1], "pages", 320);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
--fcc-editable-region--
  
--fcc-editable-region--
];
```
