---
id: 68225ea2f1aae301fa4df0f4
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

بينما تعمل مصفوفة النتائج، يجب أن تكون النتيجة النهائية المطلوبة سلسلة نصية تسرد جميع الكتب مع عنوان ومؤلف.

هنا تأتي طريقة `join()`. في الدروس السابقة، تعلمت أن طريقة `join()` تُستخدم لإرجاع سلسلة نصية جديدة تحتوي على جميع عناصر المصفوفة مدموجة في سلسلة واحدة، مع فاصل محدد بين كل عنصر.

إليك تذكيرًا:

```js
const developers = ["Naomi", "Tom", "Jessica"];

const teamList = developers.join("\n");
console.log(teamList); 
// Naomi
// Tom
// Jessica
```

الفاصل في هذه الحالة هو `\n` الذي يمثل حرف السطر الجديد.

اربط طريقة `join()` مع `\n` كفاصل إلى `map()`. الآن يجب أن ترى سلسلة نصية في وحدة التحكم بدلًا من مصفوفة النتائج.

# --hints--

يجب أن تُرجع دالة `getBookInformation` سلسلة نصية.

```js
assert.isString(getBookInformation(library));
```

يجب أن تُرجع دالة `getBookInformation` سلسلة نصية واحدة تمثل النتيجة النهائية.

```js
const testLibrary = [
  {
    title: "Title A",
    author: "Author A",
    about: "About A",
    pages: 320,
  },
  {
    title: "Title B",
    author: "Author B",
    about: "About B",
    pages: 320,
  },
  {
    title: "Title C",
    author: "Author C",
    about: "About C",
    pages: 304,
  },
];

const expected = `Title A by Author A
Title B by Author B
Title C by Author C`;

assert.strictEqual(getBookInformation(testLibrary), expected);
```

# --seed--

## --seed-contents--

```js
const library = [
  {
    title: 'Your Next Five Moves: Master the Art of Business Strategy',
    author: 'Patrick Bet-David and Greg Dinkin',
    about: 'A book on how to plan ahead',
    pages: 320,
  },
  {
    title: 'Atomic Habits',
    author: 'James Clear',
    about: 'A practical book about discarding bad habits and building good ones',
    pages: 320,
  },
  {
    title: 'Choose Your Enemies Wisely: Business Planning for the Audacious Few',
    author: 'Patrick Bet-David',
    about:
      "A book that emphasizes the importance of identifying and understanding one's adversaries to succeed in the business world",
    pages: 304,
  },
  {
    title: 'The Embedded Entrepreneur',
    author: 'Arvid Kahl',
    about: 'A book focusing on how to build an audience-driven business',
    pages: 308,
  },
  {
    title: 'How to Be a Coffee Bean: 111 Life-Changing Ways to Create Positive Change',
    author: 'Jon Gordon',
    about: 'A book about effective ways to lead a coffee bean lifestyle',
    pages: 256,
  },
  {
    title: 'The Creative Mindset: Mastering the Six Skills That Empower Innovation',
    author: 'Jeff DeGraff and Staney DeGraff',
    about: 'A book on how to develop creativity and  innovation skills',
    pages: 168,
  },
  {
    title: 'Rich Dad Poor Dad',
    author: 'Robert Kiyosaki and Sharon Lechter',
    about: 'A book about financial literacy, financial independence, and building wealth. ',
    pages: 336,
  },
  {
    title: 'Zero to Sold',
    author: 'Arvid Kahl',
    about: 'A book on how to bootstrap a business',
    pages: 500,
  },
];

console.log("Books in the Library:\n");

function getBookInformation(catalog) {
--fcc-editable-region--
  return catalog.map(book => `${book.title} by ${book.author}`)
--fcc-editable-region--
}

console.log(getBookInformation(library));
```
