---
id: afcc8d540bea9ea2669306b6
title: بناء دالة لتكرار سلسلة نصية
challengeType: 26
dashedName: build-a-string-repeating-function
---

# --description--

في هذا المختبر، ستنشئ دالة تكرر سلسلة نصية معينة عددًا محددًا من المرات. لغرض هذا المختبر، لا تستخدم الطريقة المدمجة `.repeat()`.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة باسم `repeatStringNumTimes` تأخذ معلمتين: سلسلة نصية ورقم.
2. يجب أن تُرجع الدالة السلسلة النصية مكررة العدد المحدد من المرات.
3. إذا كان الرقم أقل من أو يساوي صفرًا، يجب أن تُرجع الدالة سلسلة نصية فارغة.

# --hints--

يجب أن تنشئ دالة باسم `repeatStringNumTimes`.

```js
assert.isFunction(repeatStringNumTimes);
```

يجب أن تأخذ `repeatStringNumTimes` معلمتين.

```js
assert.lengthOf(repeatStringNumTimes, 2);
```

يجب أن تُرجع الدالة `repeatStringNumTimes` دائمًا سلسلة نصية.

```js
assert.isString(repeatStringNumTimes('hello', 3));
```

يجب أن تُرجع `repeatStringNumTimes("*", 3)` السلسلة النصية `***`.

```js
assert.strictEqual(repeatStringNumTimes('*', 3), '***');
```

يجب أن تُرجع `repeatStringNumTimes("abc", 3)` السلسلة النصية `abcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 3), 'abcabcabc');
```

يجب أن تُرجع `repeatStringNumTimes("abc", 4)` السلسلة النصية `abcabcabcabc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 4), 'abcabcabcabc');
```

يجب أن تُرجع `repeatStringNumTimes("abc", 1)` السلسلة النصية `abc`.

```js
assert.strictEqual(repeatStringNumTimes('abc', 1), 'abc');
```

يجب أن تُرجع `repeatStringNumTimes("*", 8)` السلسلة النصية `********`.

```js
assert.strictEqual(repeatStringNumTimes('*', 8), '********');
```

يجب أن تُرجع `repeatStringNumTimes("abc", -2)` سلسلة نصية فارغة (`""`).

```js
assert.isEmpty(repeatStringNumTimes('abc', -2));
```

يجب أن تُرجع `repeatStringNumTimes("abc", 0)` السلسلة `""`.

```js
assert.isEmpty(repeatStringNumTimes('abc', 0));
```

يجب عدم استخدام الطريقة المدمجة `repeat()`.

```js
assert.notMatch(__helpers.removeJSComments(code), /\.repeat/g);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function repeatStringNumTimes(str, num) {
  if (num < 1) return '';
  return num === 1 ? str : str + repeatStringNumTimes(str, num - 1);
}

repeatStringNumTimes('abc', 3);
```
