---
id: 694a9ff65d4472dd451d8255
title: 步驟 1
challengeType: 20
dashedName: step-1
---

# --description--

在這個工作坊中，你將透過建置一個電影票訂票計算器來練習如何在 Python 中使用布林值和條件述語。

首先宣告用來儲存使用者和電影放映資訊的變數。

如你在先前課程中所學，變數是使用指定運算子（`=`）以這種方式指定的：

```py
# Assigns the value 10 to variable x
x = 10
```

創建一個名為 `base_price` 的變數來儲存電影票的基本價格，並將其值設為 `15`。創建另一個名為 `age` 的變數來儲存使用者的年齡，並將其值設為 `21`。

# --hints--

你應該宣告一個名為 `base_price` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

你應該將整數值 `15` 指定給 `base_price`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

你應該宣告一個名為 `age` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

你應該將整數值 `21` 指定給 `age`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
