---
id: 694cc1b515976889afd4cded
title: 步驟 2
challengeType: 20
dashedName: step-2
---

# --description--

現在你需要儲存一些關於電影票的字串值。你可以用這種方式在 Python 中儲存字串：

```py
name = 'Alice'
```

創建一個名為 `seat_type` 的變數來儲存使用者所選擇的座位型別，並將其值設定為字串 `Gold`。創建另一個名為 `show_time` 的變數來儲存電影的放映時間，並將其值設定為字串 `Evening`。

請記得用單引號或雙引號將兩個值都包起來。

# --hints--

你應該宣告一個名為 `seat_type` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

你應該將字串 `Gold` 指定給 `seat_type`。請記得用單引號或雙引號將該值包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

你應該宣告一個名為 `show_time` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

你應該將字串 `Evening` 指定給 `show_time`。請記得用單引號或雙引號將該值包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
