---
id: 6958dd85587e10c99162dd81
title: 步驟 3
challengeType: 20
dashedName: step-3
---

# --description--

如你在先前課程中所學，在 Python 中，`if` 述語可用來根據條件是否為真來執行程式碼。

`if` 述語由 `if` 關鍵字組成，後面接著一個條件和冒號。當條件為真時要執行的程式碼，必須縮排，稱為 `if` 述語的主體。

```py
if condition:
    # Code to execute if condition is True
```

在此步驟中，你將根據使用者的年齡檢查他們是否有資格預訂電影票。

創建一個 `if` 述語來檢查 `age` 是否大於 `17`。在 `if` 述語的主體內，列印 `User is eligible to book a ticket`。只有當使用者的年齡大於 `17` 時，才會列印該訊息。

請記得縮排 `if` 述語的主體，並在 `print()` 呼叫中用單引號或雙引號將訊息包圍。

# --hints--

你應該有一個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

`if` 述語條件應該是 `age > 17`。

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

你應該在你的 `if` 述語內有一個 `print()` 呼叫。記得縮排 `if` 述語的主體。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

你應該在你的 `if` 述語中列印 `User is eligible to book a ticket`。請記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
