---
id: 6958dfa0c74b237e1b56ce04
title: 第 7 步
challengeType: 20
dashedName: step-7
---

# --description--

如果使用者是成員，則該使用者有資格獲得會員折扣。

創建一個名為 `discount` 的變數並將其值設為 `0`。這將儲存使用者在電影票上獲得的折扣。

# --hints--

你應該宣告一個名為 `discount` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

你應該將整數值 `0` 指定給 `discount`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
