---
id: 6958e0ecc121b0ae3dd9c5e0
title: 第 9 步
challengeType: 20
dashedName: step-9
---

# --description--

你也應該處理使用者不符合會員折扣資格的情況。

為你的 `if` 述語添加一個 `else` 語句，並在 `else` 主體內列印 `User does not qualify for membership discount`。

你可以這樣列印一個訊息，後面接著一個變數：

```py
print('Total:', total)
```

你也想要顯示更新後的 `discount` 值。在 `if...else` 述語下方，使用 `print()` 呼叫來顯示一則訊息，該訊息顯示 `Discount:`，後面接著更新後的 `discount` 值。然後，在終端機中檢查輸出。

# --hints--

你應該有一個 `else` 語句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

你應該在你的 `else` 語句中列印 `User does not qualify for membership discount`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

你應該在你的 `if...else` 述語下方放置 `print('Discount:', discount)`。

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
