---
id: 6958de0d773289a89515a6b8
title: 步驟 4
challengeType: 20
dashedName: step-4
---

# --description--

現在你將根據使用者的年齡檢查是否允許預訂晚場演出。

創建一個 `if` 述語來檢查 `age` 是否大於或相等於 `21`。在 `if` 述語的主體內，列印 `User is eligible for Evening shows`。

# --hints--

你的程式碼中應該有兩個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

你的第二個 `if` 述語條件應該是 `age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

你應該在你的新 `if` 述語內列印 `User is eligible for Evening shows`。記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
