---
id: 6958e2377cacbc46a7f7383a
title: 步驟 13
challengeType: 20
dashedName: step-13
---

# --description--

現在，為你的 `if` 述語添加一個 `else` 語句，並在 `else` 主體內列印 `No extra charges will be applied`。這將僅在額外費用條件為假時列印該訊息。

然後，在 `else` 語句下方，使用 `print()` 呼叫來顯示一則訊息，內容為 `Extra charges:`，後面接著更新後的 `extra_charges` 值，並在終端機中檢查輸出。

# --hints--

你應該有一個 `else` 語句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

你應該在你的 `else` 語句中列印 `No extra charges will be applied`。記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

你應該在你的 `if...else` 述語下方放置 `print('Extra charges:', extra_charges)`。

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
