---
id: 6958e1875e592b58e8852757
title: 步驟 12
challengeType: 20
dashedName: step-12
---

# --description--

現在創建一個名為 `extra_charges` 的變數並將其設為 `0`。這將表述週末電影票需額外收取的費用。

然後，創建一個 `if` 述語來檢查 `is_weekend` 是否為真。在 `if` 述語的主體內，將 `extra_charges` 的值更新為 `2`，並在終端機中列印 `Extra charges will be applied`。

# --hints--

你應該宣告一個名為 `extra_charges` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

你應該將整數值 `0` 指定給 `extra_charges`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

你應該有第四個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

你的第四個 `if` 述語條件應該是 `is_weekend`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

你應該在你的新 `if` 述語中，將整數值 `2` 指定給 `extra_charges` 變數。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

你應該在你的新 `if` 述語中列印 `Extra charges will be applied`。請記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
