---
id: 6958e59845777391b86d7331
title: 步驟 18
challengeType: 20
dashedName: step-18
---

# --description--

在 Python 中，`if` 述語也可以放在另一個 `if` 述語的主體內。這稱為巢狀 `if` 述語。

巢狀的 `if` 述語允許你在第一個條件已經滿足後，檢查額外的條件。只有當外層的 `if` 條件為真時，內層的 `if` 述語才會執行。

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

現在你將根據使用者所選擇的座位型別計算服務費用。

仍然在你最後一個 `if` 述語主體內，創建一個名為 `service_charges` 的變數並將其設為 `0`。

然後，創建一個巢狀的 `if` 述語來檢查 `seat_type` 是否相等於 `Premium`。在 `if` 述語的主體內，將 `service_charges` 的值更新為 `5`。

# --hints--

你應該在最後一個 `if` 述語區塊內宣告一個名為 `service_charges` 的變數。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

你應該將整數值 `0` 指定給 `service_charges`。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

你應該在最後一個 `if` 述語的主體內有一個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

新的 `if` 述語條件應該是 `seat_type == 'Premium'`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

你應該在你的新 `if` 述語內將整數值 `5` 指定給 `service_charges` 變數。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
