---
id: 696099abb6acc539b27055a1
title: 步驟 5
challengeType: 20
dashedName: step-5
---

# --description--

在前一步驟中，你使用 `if` 述語檢查了一個條件。當條件為假時，`else` 語句用來控制代碼。以下是 `if...else` 述語的語法：

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

現在，為你的 `if` 語句添加一個 `else` 語句，並在 `else` 主體內列印 `User is not eligible for Evening shows`。只有當 `if` 語句中的條件為假時，才會列印此內容。

# --hints--

你應該有一個 `else` 語句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

你應該在你的 `else` 語句中列印 `User is not eligible for Evening shows`。記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
