---
id: 6960a5165f7efe52a0c0aa39
title: 第 10 步
challengeType: 20
dashedName: step-10
---

# --description--

在 Python 中，`and` 運算子用於檢查多個條件是否為真。以下是你如何在 `if` 述語中使用它來結合兩個條件：

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

會員折扣應該只適用於其 `age` 大於或相等於 `21` 的成員。

使用 `and` 運算子將你 `if` 述語的現有條件與另一個檢查 `age` 是否大於或相等於 `21` 的條件結合。

# --hints--

你的第三個 `if` 述語條件應該是 `is_member and age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
