---
id: 6960b6a6ea06e27dcc7eef37
title: 第 15 步
challengeType: 20
dashedName: step-15
---

# --description--

現在你將檢查使用者是否符合訂購電影票的條件。年齡為 `21` 歲或以上的使用者總是可以在沒有任何限制的情況下訂票。

創建一個 `if` 述語來檢查 `age` 是否大於或相等於 `21`。在 `if` 述語的主體內，將 `Ticket booking condition satisfied` 列印到終端機。

然後，為你的 `if` 述語添加一個 `else` 語句，並在 `else` 主體內列印 `Ticket booking failed due to restrictions`。

# --hints--

你應該有第五個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

你的第五個 `if` 述語條件應該是 `age >= 21`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

你應該在你的新 `if` 述語內列印 `Ticket booking condition satisfied`。記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

你的新 `if` 述語應該有一個 `else` 語句。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

你應該在你的 `else` 語句中列印 `Ticket booking failed due to restrictions`。請記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
