---
id: 6960b308bf41cc39c9b4b438
title: 第 14 步
challengeType: 20
dashedName: step-14
---

# --description--

在 Python 中，`or` 運算子用於檢查兩個條件中是否至少有一個為真。以下是你如何在 `if` 述語中使用它：

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

如果演出是在晚上，則應該也適用額外收費。使用 `or` 運算子將你 `if` 述語中現有的條件與第二個條件結合，該條件檢查 `show_time` 是否相等於字串 `Evening`。

# --hints--

你最後的 `if` 述語條件應該是 `is_weekend or show_time == 'Evening'`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
