---
id: 696099c1ac72a4e585de1d5b
title: 步驟 6
challengeType: 20
dashedName: step-6
---

# --description--

有些資訊只能是真或假。正如你在先前課程中學到的，這可以用布林值來表述。

創建一個名為 `is_member` 的變數來顯示使用者是否為成員，並將其值設定為 `True`。

在 `is_member` 變數下方創建另一個名為 `is_weekend` 的變數，以顯示電影場次是否在週末，並將其值設定為 `False`。不要用引號包住該值。

# --hints--

你應該宣告一個名為 `is_member` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

你應該將布林值 `True` 指定給 `is_member`。不要用引號包住該值。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

你應該宣告一個名為 `is_weekend` 的變數。

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

你應該將布林值 `False` 指定給 `is_weekend`。不要用引號包住該值。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
