---
id: 6960a848aacf859f7bb15447
title: 第 11 步
challengeType: 20
dashedName: step-11
---

# --description--

現在將 `is_member` 變數的值改為 `False`，因為使用者不是成員。

之後，你會看到 `discount` 值現在保持為 `0`，因為折扣必須同時滿足兩個條件才會生效。

# --hints--

你應該將 `is_member` 變數的值更新為 `False`。不要用引號包住該值。

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
