---
id: 6960ba6e6c95ed4ad896c6f2
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

訂票規則還有一個例外。年齡介於 `18` 和 `21` 之間的使用者，如果是成員，可以訂購晚場演出。

當在 `if` 述語中使用多個邏輯運算子時，使用 `and` 連接的條件會在使用 `or` 連接的條件之前被評估。Python 中使用小括弧 `()` 來將條件群組並控制它們被評估的順序。

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

現在，使用 `or` 運算子為你現有的 `if` 述語添加另一個條件，以檢查 `is_member` 是否為真值。使用小括弧 `()` 將 `show_time != 'Evening'` 和 `is_member` 條件一起分組，如上述範例所示。

# --hints--

你最後的 `if` 述語條件應該是 `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
