---
id: 696cb889fc4fdeb337e52864
title: 第 8 步
challengeType: 20
dashedName: step-8
---

# --description--

在 Python 中，`if` 條件不必明確地將值與 `True` 或 `False` 比較。相反地，它們依賴於值的真值性。真值性值在布林值背景關係中會被評估為 `True`，例如 `if` 述語的條件。這些包括非零數字、非空字串以及布林值 `True`。

```py
if value:
   print('value is truthy')
```

創建一個 `if` 述語來檢查 `is_member` 是否為真值。在 `if` 述語的主體內，將 `discount` 的值更新為 `3`，並在終端機列印 `User qualifies for membership discount`。

# --hints--

你的程式碼中應該有三個 `if` 述語。

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

你的第三個 `if` 述語條件應該是 `is_member`。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

你應該在你的新 `if` 述語中，將整數值 `3` 指定給 `discount` 變數。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

你應該在新的 `if` 述語中列印 `User qualifies for membership discount`。記得用單引號或雙引號將訊息包起來。

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
