---
id: 696cbb3985e23420cd45c58f
title: 第 20 步
challengeType: 20
dashedName: step-20
---

# --description--

`if...elif...else` 述語用於依序檢查多個條件。

```py
if condition1:
   # Code to execute if condition1 is True
elif condition2:
   # Code to execute if condition1 is False and condition2 is True
else:
   # Code to execute if all conditions are False
```

為你的 `if...else` 述語添加一個 `elif` 述語並檢查 `seat_type` 是否相等於 `Gold`。在 `elif` 述語的主體內，將 `service_charges` 的值更新為 `3`。

在 `if...elif..else` 述語下方，使用 `print()` 呼叫來顯示一則訊息，內容為 `Service charges:`，後面接著更新後的 `service_charges` 值。然後，在終端機中檢查輸出。

# --hints--

你應該有一個 `elif` 述語。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1]
assert not _cond.is_empty()
`)) })
```

你的 `elif` 述語條件應該是 `seat_type == 'Gold'`。

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1] 
assert _cond.is_equivalent("seat_type == 'Gold'")
`)) })
```

你應該在你的 `elif` 區塊內，將整數值 `3` 指定給 `service_charges` 變數。

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[1].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 3")
`)) })
```

你應該在你的 `if...elif...else` 述語下方放置 `print('Service charges:', service_charges)`。

```js
({ test: () => (runPython(`
if_stmt = """if seat_type == 'Premium':
    service_charges = 5
elif seat_type == 'Gold':
    service_charges = 3
else:
    service_charges = 1"""
print_call = "print('Service charges:', service_charges)"
assert _Node(_code).find_ifs()[4].find_bodies()[0].is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

    service_charges = 0
    if seat_type == 'Premium':
        service_charges = 5
--fcc-editable-region--
    
    else:
        service_charges = 1
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
