---
id: 6995d0325fd7da40854fac40
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في درس سابق، تعلمت عن البحث الثنائي وكيف يختلف عن البحث الخطي.

في هذه الورشة، ستطبق ذلك من خلال بناء خوارزمية بحث ثنائي.

ابدأ بإنشاء دالة `binarySearch` تحتوي على معلمتين `searchList` و `value`.

`searchList` هي المصفوفة التي ستبحث الدالة داخلها عن العناصر، و`value` هو الهدف الذي تحاول الدالة إيجاده داخل `searchList`.

# --hints--

يجب أن تنشئ دالة `binarySearch`.

```js
assert.isFunction(binarySearch);
```

يجب أن تحتوي دالة `binarySearch` على معلمتين `searchList` و `value`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
