---
id: 6995d0325fd7da40854fac40
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في درس سابق، تعلمت عن البحث الثنائي وكيف يختلف عن البحث الخطي.

في هذه الورشة، ستطبق ذلك من خلال بناء خوارزمية بحث ثنائي.

ابدأ بإنشاء دالة `binarySearch` تحتوي على معلمتين `searchList` و `value`.

`searchList` هي المصفوفة التي ستبحث الدالة داخلها عن العناصر، و`value` هو الهدف الذي تحاول الدالة إيجاده داخل `searchList`.

# --hints--

يجب أن تنشئ دالة `binarySearch`.

```js
assert.isFunction(binarySearch);
```

يجب أن تحتوي دالة `binarySearch` على معلمتين `searchList` و `value`.

```js
const regex = __helpers.functionRegex('binarySearch', ['searchList', 'value']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
