---
id: 6996c60281feca6718cc15c1
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

داخل دالة `binarySearch`، عرّف متغيرًا باسم `pathToTarget` وابدأ بتهيئته كمصفوفة فارغة.

# --hints--

يجب أن يكون لديك متغير `pathToTarget` مضبوطًا على مصفوفة فارغة `[]` داخل دالة `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
