---
id: 6996d1032feed88e2957dc8e
title: الخطوة 15
challengeType: 1
dashedName: step-15
---

# --description--

يمكنك أن ترى أن الاستدعاء الأخير يطبع مصفوفة فارغة. لتوضيح أن القيمة غير موجودة في المصفوفة التي تم البحث فيها، حدّث جملة `return` لتُعيد مصفوفة تحتوي على عنصرين، مصفوفة فارغة، والسلسلة النصية `Value not found`.

# --hints--

يجب تحديث جملة `return` لتُعيد `[], 'Value not found'`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /return\s*\[\s*\[\s*\]\s*,\s*('|"|`)Value not found\1\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
--fcc-editable-region--
  return [];
--fcc-editable-region--
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
