---
id: 6996d0a2c2ac55041145a13f
title: الخطوة 14
challengeType: 1
dashedName: step-14
---

# --description--

هل لاحظت أن الاستدعاء الثاني للدالة يحتوي الآن على القيم `3, 5, 4` في القائمة؟ هذا يدل على أن البحث الثنائي يعمل كما هو مقصود.

هكذا حدث ذلك: الخوارزمية فحصت أولًا `3` كمنتصف القائمة الأولية. بما أن `4` أكبر من `3`، تحوّل البحث إلى النصف الأيمن. ثم فحصت `5` كمنتصف جديد. وبما أن `4` أقل من `5`، تحرك البحث إلى اليسار، محددًا في النهاية `4` كمنتصف النطاق النهائي.

لاختبار الدالة مرة أخرى، استدعها مع `[1, 3, 5, 9, 14, 22], 10` وسجل النتيجة في وحدة التحكم. هذه حالة لن يتم فيها العثور على القيمة.

# --hints--

يجب أن تستدعي الدالة باستخدام `binarySearch([1, 3, 5, 9, 14, 22], 10)` وتسجلها في وحدة التحكم.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
