---
id: 6996c796385b6ea4e222d5e8
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

لتبدأ مع الحلقة، عليك إيجاد فهرس المنتصف لمساحة البحث الحالية.

يمكنك الحصول على فهرس المنتصف بحساب متوسط المتغيرين `low` و`high`. يجب عليك استخدام القسمة الأرضية للحصول على المتوسط بعد جمع القيمتين، ليكون الناتج دائمًا عددًا صحيحًا مقربًا للأسفل، لأن الفهارس يجب أن تكون أعدادًا صحيحة.

داخل حلقة `while`، أنشئ متغيرًا `mid` يساوي متوسط `low` و`high`.

# --hints--

داخل حلقة while يجب أن تنشئ متغيرًا `mid` يساوي `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
