---
id: 6996ca21a74104721f2c7e5d
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

تحتاج إلى اختبار الأمور لتفهم تدفق الخوارزمية في هذه المرحلة الأولية.

للقيام بذلك، يجب عليك أولاً الخروج من الحلقة. ذلك لأن التنفيذ الحالي يسمح بتكرار واحد فقط، فإذا لم يتحقق الشرط، ستحدث حلقة لا نهائية.

بعد جملة `if` مباشرة، استخدم الكلمة المفتاحية `break` للخروج من حلقة `while`. ثم، بعد حلقة `while`، أرجع مصفوفة فارغة للدلالة على أن القيمة لم تُعثر عليها.

# --hints--

يجب عليك استخدام الكلمة المفتاحية `break` للخروج من حلقة `while` بعد جملة `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

يجب أن ترجع مصفوفة فارغة بعد حلقة `while`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
