---
id: 6996cb5869acf58de751ae81
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

تذكّر أن الدالة تأخذ مصفوفة تسمى `searchList` والقيمة `value` التي تبحث عنها.

الآن، استدعِ الدالة مع `binarySearch([1, 2, 3, 4, 5], 3)` و`binarySearch([1, 2, 3, 4, 5, 9], 4)` وسجل كلا البيانين في وحدة التحكم.

# --hints--

يجب أن تستدعي الدالة مع `binarySearch([1, 2, 3, 4, 5], 3)` وتسجلها في وحدة التحكم.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

يجب أن تستدعي الدالة مع `binarySearch([1, 2, 3, 4, 5, 9], 4)` وتسجلها في وحدة التحكم.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
