---
id: 6996c7363a4a437839f1f15d
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

الآن بعد أن عرّفت حدود البحث، حان الوقت لإنشاء الحلقة التي ستنفذ خوارزمية البحث الثنائي.

يعمل البحث الثنائي عن طريق تضييق مساحة البحث بشكل متكرر. تستمر هذه العملية طالما هناك نطاق صالح من العناصر للفحص.

يمكنك التعبير عن ذلك باستخدام حلقة `while` تستمر طالما أن مؤشر `low` أقل من أو يساوي مؤشر `high`.

إذا أصبح `low` أكبر من `high`، فهذا يعني أن مساحة البحث أصبحت فارغة، وأن `value` غير موجود في المصفوفة.

داخل دالة `binarySearch`، أنشئ حلقة `while` بشرط يتحقق مما إذا كان `low` أقل من أو يساوي `high`.

# --hints--

يجب أن تنشئ حلقة `while` تعمل طالما أن `low` أقل من أو يساوي `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
