---
id: 6996c8375f4301004459ee85
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

الآن بعد أن حسبت فهرس `mid`، تحتاج فعليًا إلى استرجاع القيمة من `searchList` عند ذلك الفهرس. هذه القيمة هي التي ستقارنها مع معلمة `value`، وهي الهدف الذي تبحث عنه.

أنشئ متغيرًا باسم `valueAtMiddle` وقم بتعيين العنصر من `searchList` الموجود عند فهرس `mid` له.

# --hints--

يجب أن تنشئ متغيرًا `valueAtMiddle` وتعيّنه إلى `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
