---
id: 6996ccadf9048c79b857bb15
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

يمكنك أن ترى أن الاستدعاء الأول يُرجع `[3]`. ذلك لأن `3` هو القيمة الوسطى في `[1, 2, 3, 4, 5]`.

من ناحية أخرى، الاستدعاء الثاني يُرجع مصفوفة فارغة. يحدث هذا لأن الحلقة بعد التحقق من القيمة الوسطى تفتقر حاليًا إلى المنطق الذي يضيق نطاق البحث بتحديث المتغيرين `low` أو `high`. هي فقط تقوم بفحص نقطة منتصف واحدة ثم تنتهي دون متابعة البحث.

للسماح لخوارزمية البحث الثنائي بمواصلة تضييق نطاق البحث، أضف جملة `else if` تتحقق مما إذا كانت `value` أكبر من `valueAtMiddle`.

# --hints--

يجب أن تضيف جملة `else if` بشرط `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
--fcc-editable-region--
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
