---
id: 6996c60281feca6718cc15c1
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

داخل دالة `binarySearch`، عرّف متغيرًا باسم `pathToTarget` وابدأ بتهيئته كمصفوفة فارغة.

# --hints--

يجب إعلان متغير `pathToTarget` وتعيينه إلى مصفوفة فارغة داخل دالة `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
