---
id: 6996c95e5bede03bdf2937a2
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

الخطوة التالية هي إنشاء شرط يتحقق مما إذا كانت القيمة المستهدفة `value` في المنتصف.

داخل حلقة `while`، أنشئ جملة `if` تتحقق مما إذا كانت القيمة المستهدفة `value` تساوي `valueAtMiddle`. إذا كانت كذلك، أرجع المتغير `pathToTarget`.

# --hints--

يجب أن تنشئ جملة `if` بالشرط `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

يجب أن تُرجع المتغير `pathToTarget` داخل جملة الـ `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
