---
id: 6996c8f94fe6b250f11a2a1e
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

الآن، أضف `valueAtMiddle` إلى مصفوفة `pathToTarget` لتتبع الخطوات التي تم اتخاذها أثناء البحث، بغض النظر عن كونها الهدف أم لا. يمكنك فعل ذلك باستخدام طريقة `push()`.

# --hints--

يجب عليك استخدام طريقة `push()` لإضافة `valueAtMiddle` إلى مصفوفة `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
