---
id: 6996c675b599baa7b103192d
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في البحث الثنائي، يوجد أدنى فهرس ممكن وأعلى فهرس ممكن.

يمثل أدنى فهرس ممكن الحد الأيسر في الجزء الحالي من القيم التي يتم البحث ضمنها، في هذه الحالة، `searchList`. أما أعلى فهرس ممكن، فيمثل الحد الأيمن في القيم التي يتم البحث ضمنها.

لأخذ هذين الحدين في الاعتبار، يجب أن تراعي المصفوفة كاملة التي يتم البحث فيها. لذلك، داخل دالة `binarySearch`، عرّف متغيرًا باسم `low` بقيمة `0`، ومتغيرًا باسم `high` بقيمة تمثل آخر فهرس في المصفوفة التي يتم البحث فيها.

# --hints--

يجب أن يكون لديك متغير باسم `low` مضبوط على القيمة `0` داخل دالة `binarySearch` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

يجب أن يكون لديك متغير باسم `high` مضبوط على القيمة `searchList.length - 1` داخل دالة `binarySearch` الخاصة بك.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
