---
id: 6996c675b599baa7b103192d
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

في البحث الثنائي، يوجد أدنى فهرس ممكن وأعلى فهرس ممكن.

يمثل أدنى فهرس ممكن الحد الأيسر في الجزء الحالي من القيم التي يتم البحث ضمنها، في هذه الحالة، `searchList`. أما أعلى فهرس ممكن، فيمثل الحد الأيمن في القيم التي يتم البحث ضمنها.

لأخذ هذين في الاعتبار، عليك أن تراعي المصفوفة كاملة التي يتم البحث فيها. لذا، داخل دالة `binarySearch`، أعلن عن متغير `low` بقيمة `0`، ومتغير `high` بقيمة تمثل آخر فهرس في المصفوفة التي يتم البحث فيها.

# --hints--

يجب إعلان متغير باسم `low` بقيمة `0` داخل دالة `binarySearch` الخاصة بك.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

يجب إعلان متغير باسم `high` بقيمة `searchList.length - 1` داخل دالة `binarySearch` الخاصة بك.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
