---
id: 6996cfb6c01e9a802add38be
title: الخطوة 13
challengeType: 1
dashedName: step-13
---

# --description--

أخيرًا، إذا لم يُعثر على `value` في المنتصف ولم يكن أكبر من `valueAtMiddle`، إذن يجب أن يكون `value` أقل من `valueAtMiddle`، ويجب أن يكون على اليسار. ستتعامل كتلة `else` النهائية مع هذا.

لذا، أنهِ الحلقة بحذف الكلمة المفتاحية `break` وإضافة كتلة `else`. بداخلها، حدّث المتغير `high` بطرح `1` من `mid`.

# --hints--

يجب استبدال الكلمة المفتاحية `break` بكتلة `else` تقوم بتحديث المتغير `high` إلى `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
--fcc-editable-region--      
    } 
    break;
--fcc-editable-region--
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
