---
id: 6551f4259d66026ee28cb205
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Ви також можете отримувати доступ до символів рядка, починаючи з кінця рядка. Останній символ має індекс `-1`, передостанній — `-2` і так далі.

Тепер змініть свій існуючий виклик `print()`, щоб вивести останній символ у вашому рядку.

# --hints--

Ви все ще повинні викликати функцію `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Ви повинні передати `text[-1]` у функцію `print()`, включивши його між круглими дужками. Зверніть увагу, щоб розмістити виклик функції на початку лінії.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[6])
--fcc-editable-region--
```
