---
id: 6551ef504d91936d2d4e54f8
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

*Аргумент* — це об’єкт або вираз, який передається функції — додається між відкриваючими та закриваючими круглими дужками — під час її виклику:

```py
greet = 'Hello!'
print(greet)
```

Код у наведеному прикладі виведе рядок `'Hello!'`, який є значенням змінної `greet`, переданої функції `print()` як аргумент.

Виведіть на екран вашу змінну `text`, передавши змінну `text` як аргумент функції `print()`.

# --hints--

Ви повинні передати `text` функції `print()`, додавши ім’я цієї змінної між відкриваючими та закриваючими круглими дужками. Зверніть увагу, що виклик функції має бути на початку рядка.

```js
assert.match(code, /^print\s*\(\s*text\s*\)/m)
```

У вашому коді є помилка синтаксису та/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print()
--fcc-editable-region--
```
