---
id: 6551fe3b1df7c9740f13f270
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Ключові аспекти іменування змінних у Python:

- Деякі слова є зарезервованими ключовими словами (наприклад, `for`, `while`, `True`). Вони мають особливе значення в Python, тому їх не можна використовувати для імен змінних.
- Імена змінних не можуть починатися з числа, і вони можуть містити лише буквено-цифрові символи або підкреслення.
- Імена змінних чутливі до регістру, тобто `my_var` відрізняється від `my_Var` і `MY_VAR`.
- Нарешті, загальною конвенцією є написання імен змінних у `snake_case`, де кожен пробіл замінений символом підкреслення, а слова написані в нижньому регістрі.

Видаліть обидва виклики `print()` і оголосіть іншу змінну з ім’ям `alphabet`. Призначте цій змінній рядок `'abcdefghijklmnopqrstuvwxyz'`.

# --hints--

У вашому коді не повинно бути `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

У вашому коді не повинно бути `print(type(shift))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/)
```

Ви повинні оголосити змінну з ім’ям `alphabet`. Зверніть увагу, що ім’я змінної має бути на початку лінії.

```js
assert.match(code, /^alphabet\s*=/m)
```

Ви повинні призначити рядок `'abcdefghijklmnopqrstuvwxyz'` змінній `alphabet`. Пам’ятайте, що рядок потрібно взяти в одинарні або подвійні лапки.

```js
assert.match(code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*(#.*)?$/m)
```

Ваш код містить недійсний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(type(shift))
--fcc-editable-region--
```
