---
id: 6551f3bfed58796ebb1268dc
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

Кожен символ рядка можна отримати за числовим індексом. Підрахунок індексів починається з нуля. Тож перший символ рядка має індекс `0`. Наприклад, у рядку `'Hello World'` `'H'` знаходиться на індексі `0`, `'e'` — на індексі `1` і так далі.

До кожного символу рядка можна отримати доступ, використовуючи дужкову нотацію. Потрібно написати ім'я змінної, після якого йдуть квадратні дужки, а між дужками вказати індекс символу:

```py
text = 'Hello World'
r = text[8]
```

Тепер, замість того щоб виводити `text`, виведіть лише символ на індексі `6`.

# --hints--

Ви все ще повинні викликати функцію `print()`.

```js
assert.match(code, /print\s*\(.*\)/)
```

Ви повинні передати `text[6]` у функцію `print()`, включивши його між круглими дужками. Зверніть увагу, що виклик функції має бути на початку лінії.

```js
assert.match(code, /^print\s*\(\s*text\s*\[\s*6\s*\]\s*\)/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text)
--fcc-editable-region--
```
