---
id: 6551ecf597038c36637f3759
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

Змінні є невід’ємною частиною Python і будь-якої мови програмування. Змінна — це ім’я, яке є референсом або вказує на об’єкт. Ви можете оголосити змінну, написавши ім’я змінної зліва від оператора присвоєння `=` і вказавши значення, яке потрібно призначити цій змінній, справа від оператора присвоєння:

```py
variable_name = value
```

Створіть змінну з ім’ям `number` і призначте значення `5` вашій новій змінній.

# --hints--

Ви повинні оголосити змінну з ім’ям `number`. Зверніть увагу, що ім’я змінної має бути на початку лінії.

```js
assert.match(code, /^number\s*=/m)
```

Ви повинні призначити значення `5` вашій змінній `number`.

```js
assert.match(code, /^number\s*=\s*5\s*(#.*)?$/m)
```

У вашому коді є недійсний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
