---
id: 6551f5756c20146fc94f8675
title: Крок 7
challengeType: 20
dashedName: step-7
---

# --description--

Ви можете отримати доступ до числа символів у рядку за допомогою вбудованої функції `len()`.

Змініть свій існуючий виклик `print()`, передавши `len(text)` замість `text[-1]`.

# --hints--


Ви повинні викликати функцію `len()`.

```js
assert.match(code, /len\s*\(.*\)/)
```

Ви повинні передати змінну `text` у функцію `len()`, включивши її між круглими дужками.

```js
assert.match(code, /len\s*\(\s*text\s*\)/)
```

Ви повинні передати `len(text)` у функцію `print()`, включивши її між круглими дужками. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/m)
```

У вашому коді не повинно бути `print(text[-1])`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*text\s*\[\s*-\s*1\s*\]\s*\)/)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(text[-1])
--fcc-editable-region--
```
