---
id: 6551f22653e6ce6e4f1d1370
title: Крок 11
challengeType: 20
dashedName: step-11
---

# --description--

Змініть свій виклик `print(shift)`, щоб вивести тип даних вашої змінної `shift`.

# --hints--

Ви повинні зберегти свій існуючий виклик `print(type(text))`.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

Ви повинні передати `shift` у функцію `type()`.

```js
assert.match(code, /type\s*\(\s*shift\s*\)/)
```

Ви повинні передати `type(shift)` у вашу функцію `print()`. Зверніть увагу, що виклик функції має бути на початку лінії.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*shift\s*\)\s*\)/m)
```

У вашому коді не повинно бути `print(shift)`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*shift\s*\)/)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
print(shift)
--fcc-editable-region--
```
