---
id: 6551fb531ce7f773071a0b84
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Як бачите, вихідні дані при виведенні `type(text)` — це `<class 'str'>`, що означає, що ваша змінна є рядком, позначеним як `str`.

Тепер перейдіть на новий рядок і створіть іншу змінну з назвою `shift`, призначивши цій змінній значення `3`.

# --hints--

Ви повинні оголосити змінну з назвою `shift`. Зверніть увагу, що ім’я змінної має бути на початку рядка.

```js
assert.match(code, /^shift\s*=/m)
```

Ви повинні призначити значення `3` вашій змінній `shift`.

```js
assert.match(code, /^shift\s*=\s*3\s*(#.*)?$/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
--fcc-editable-region--
```
