---
id: 6551f1e8737dbe6e0cb08ea4
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер виведіть вашу нову змінну.

# --hints--

Ви не повинні видаляти свій існуючий виклик `print(type(text))`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/)
```

Ви маєте викликати функцію `print()`, передаючи змінну `shift`. Зверніть увагу, що виклик функції має бути на початку лінії.

```js
assert.match(code, /^print\s*\(\s*shift\s*\)/m)
```

У вашому коді є неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(type(text))
shift = 3
--fcc-editable-region--
```
