---
id: 6551f6f9ec6db270eb83f3db
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

Ви бачите, що `11` виведено на термінал, тому що `'Hello World'` містить 11 символів.

Ще одна корисна вбудована функція — це `type()`, яка повертає тип даних змінної. Змініть свій виклик `print()`, щоб вивести тип даних `text`.

# --hints--

Ви повинні викликати функцію `type()`.

```js
assert.match(code, /type\s*\(.*\)/)
```

Ви повинні передати `text` у функцію `type()`, включивши її між круглими дужками.

```js
assert.match(code, /type\s*\(\s*text\s*\)/)
```

Ви повинні передати `type(text)` у функцію `print()`, включивши її між круглими дужками. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*type\s*\(\s*text\s*\)\s*\)/m)
```

У вашому коді не повинно бути `print(len(text))`.

```js
const commentless_code = __helpers.python.removeComments(code)
assert.notMatch(commentless_code, /print\s*\(\s*len\s*\(\s*text\s*\)\s*\)/)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
print(len(text))
--fcc-editable-region--
```
