---
id: 65521203d7165c7b84b22ad4
title: Крок 23
challengeType: 20
dashedName: step-23
---

# --description--

Цикл дозволяє систематично проходити послідовність елементів і виконувати дії над кожним із них.

У цьому випадку ви використаєте цикл `for`. Ось як можна ітерувати над `text`:

```py
for i in text:
```

`for` — це ключове слово, що позначає тип циклу. `i` — змінна, яка послідовно приймає значення елементів у `text`. Інструкція завершується двокрапкою, `:`.

Під рядком, де ви оголосили `alphabet`, напишіть цикл `for`, щоб ітерувати над `text`. Використайте `i` як змінну циклу.

При цьому в терміналі виникає помилка. Ви дізнаєтеся про неї в наступному кроці.

# --hints--

Ви повинні використати ключове слово `for` для створення циклу. Переконайтеся, що ключове слово `for` розташоване на початку рядка, а після нього є порожній простір.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+/m)
```

Після ключового слова `i` слід написати змінну `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i/m)
```

Після `in` слід написати ключове слово `for i `. Переконайтеся, що навколо ключового слова `in` є пробіли.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+/m)
```

Після `text` слід написати `for i in `. Не забудьте додати фінальне `:`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^for\s+i\s+in\s+text\s*:\s*$/m)
```

Ваш цикл `for` має бути розміщений під рядком коду `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1\s*^for\s+i\s+in\s+text\s*:\s*$/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

--fcc-editable-region--
```
