---
id: 655261b2e1f2c197093f3993
title: Крок 41
challengeType: 20
dashedName: step-41
---

# --description--

Наразі пробіли шифруються як `'c'`. Щоб зберегти оригінальні пробіли у відкритому повідомленні, вам знадобиться умовна інструкція `if`. Вона складається з ключового слова `if`, умови та двокрапки `:`.

```py
if x != 0:
    print(x)
```

У наведеному вище прикладі умова інструкції `if` — це `x != 0`. Код `print(x)`, всередині тіла інструкції `if`, виконується лише тоді, коли умова оцінюється як `True` (у цьому прикладі це означає, що `x` відмінне від нуля).

На початку вашого циклу for замініть `print(char == ' ')` на інструкцію `if`. Умова цієї інструкції `if` має оцінюватися як `True`, якщо `char` — порожній пробіл, інакше `False`. Усередині тіла `if` виведіть рядок `'space!'`. Не забудьте відступити цей рядок.

# --hints--

У вашому коді не повинно бути `print(char == ' ')`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert.notMatch(block_body, /print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/);
```

Ви маєте замінити `print(char == ' ')` на інструкцію `if`, яка спрацьовує, коли `char == ' '`. Не використовуйте круглі дужки для умови `if` і пам’ятайте про двокрапку в кінці.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/if\s+char\s*==\s*("|')\s\1\s*:/));
```

Ви маєте вивести рядок `'space!'` всередині нової інструкції `if`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /if\s+char\s*==\s*("|')\s\3\s*/);    
assert(block_body.match(/print\s*\(\s*("|')space!\1\s*\)/));
```

Ваш код містить неправильний синтаксис і/або неправильні відступи.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    print(char == ' ')
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
