---
id: 655208d59b131e7816f18c96
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Функція `print()` виводить лише вихідні дані в консоль, але функції та методи можуть мати повернене значення, яке ви можете використовувати у своєму коді.

Тепер призначте `alphabet.find(text[0])` змінній з ім’ям `index`. Таким чином, `index` збереже значення, повернене `alphabet.find(text[0])`.

# --hints--

Ви повинні оголосити змінну з ім’ям `index`. Зверніть увагу, що ім’я змінної має бути на початку лінії.

```js
assert.match(code, /^index\s*=/m)
```

Ви повинні призначити значення, повернене `alphabet.find(text[0])`, вашій змінній `index`.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\)\s*(#.*)?$/m)
```

У вашому коді є неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find(text[0])
--fcc-editable-region--
```
