---
id: 65520f8282faf57a0db4f7fe
title: Крок 21
challengeType: 20
dashedName: step-21
---

# --description--

Як видно з вихідних даних, `'h'` знаходиться на індексі `7` у рядку `alphabet`. Тепер потрібно знайти літеру на індексі `7` плюс значення `shift`. Для цього можна використати оператор додавання, `+`, так само, як ви б його використали для математичного додавання.

Змініть вашу змінну `shifted`, щоб вона зберігала значення `alphabet` на індексі `index + shift`.

# --hints--

Ви повинні призначити `alphabet[index + shift]` вашій змінній `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*(index|shift)\s*\+\s*(?!\1)(index|shift)\s*\]\s*(#.*)?$/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
print(shifted)
--fcc-editable-region--
```
