---
id: 655209a4a27dd37873c4cac3
title: Крок 17
challengeType: 20
dashedName: step-17
---

# --description--

`.find()` повертає індекс відповідного символу всередині рядка. Якщо символ не знайдено, він повертає `-1`. Як ви бачите, перший символ у `text`, верхньому регістрі `'H'`, не знайдений, оскільки `alphabet` містить лише літери нижнього регістру.

Ви можете перетворити рядок у його еквівалент у нижньому регістрі за допомогою методу `.lower()`. Додайте ще один виклик `print()`, щоб вивести `text.lower()` і побачити вихідні дані.

# --hints--

У вашому коді має залишатися `print(index)`. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

Ви маєте вивести `text.lower()`. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
--fcc-editable-region--
```
