---
id: 655251308f31958d06cdf267
title: Крок 36
challengeType: 20
dashedName: step-36
---

# --description--

Наразі зашифрований символ оновлюється на кожній ітерації. Було б краще зберігати зашифрований рядок у новій змінній. Перед вашим `for` циклом оголосіть змінну з назвою `encrypted_text` і призначте цій змінній порожній рядок (`''`).

# --hints--

Ви повинні оголосити змінну з назвою `encrypted_text` перед вашим `for` циклом.

```js
({ test: () => assert(__userGlobals.has("encrypted_text")) })
```

Ви повинні призначити порожній рядок вашій змінній `encrypted_text`.

```js
({ test: () => assert.equal(__userGlobals.get("encrypted_text"), "") })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    new_char = alphabet[new_index]
    print('char:', char, 'new char:', new_char)
--fcc-editable-region--
```
