---
id: 6552487e689f6e895f658717
title: Крок 33
challengeType: 20
dashedName: step-33
---

# --description--

Тепер потрібно створити змінну `new_char` в кінці тіла вашого циклу. Присвойте їй значення `alphabet[new_index]`.

# --hints--

Ви маєте створити змінну `new_char` всередині вашого циклу `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_char\s*=/));
```

Ви маєте присвоїти вашій змінній `new_char` значення `alphabet[new_index]` в кінці тіла циклу.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_char\s*=\s*alphabet\s*\[\s*new_index\s*\]\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift
--fcc-editable-region--
```
