---
id: 65521fc818947e800bffe48a
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Всередині циклу `for`, перед тим як вивести поточний символ, оголосіть змінну з назвою `index` і призначте цій змінній значення, яке повертає `alphabet.find(char)`.

# --hints--

Ви повинні оголосити нову змінну з ім’ям `index` на початку вашого циклу `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/^\s*index\s*=/));
```

Ви повинні призначити `alphabet.find(char)` вашій новій змінній `index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\s*/);
assert(block_body.match(/index\s*=\s*alphabet\.find\s*\(\s*char\s*\)\s*(#.*)?$/m));
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for char in text:
    print(char)
--fcc-editable-region--
```
