---
id: 65524d2a1a253b8bb5197ae2
title: Крок 28
challengeType: 20
dashedName: step-28
---

# --description--

`find` знову повертає `-1` для літер у верхньому регістрі, а також для символу пробілу. Ви пізніше подбаєте про пробіл.

Поки що, замість ітерування над `text`, змініть цикл `for` так, щоб він ітерував над `text.lower()`.

# --hints--

Вам слід змінити свій цикл `for`, щоб ітерувати над `text.lower()` замість `text`.

```js
({test: () => assert.match(code, /^for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text:
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
