---
id: 65522255d5b9cd80f335c6f2
title: Крок 29
challengeType: 20
dashedName: step-29
---

# --description--

Наприкінці тіла вашого циклу оголосіть змінну з назвою `new_index` і призначте цій змінній значення `index + shift`.

# --hints--

Ви повинні оголосити змінну з назвою `new_index` всередині вашого циклу `for`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
assert(block_body.match(/new_index\s*=/));
```

Ви повинні призначити `index + shift` вашій новій змінній наприкінці тіла циклу `for`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_indentation, block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/new_index\s*=\s*index\s*\+\s*shift\s*$/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
--fcc-editable-region--
```
