---
id: 6552008c0d9d9075cbec9772
title: Крок 13
challengeType: 20
dashedName: step-13
---

# --description--

Ви збираєтеся використати метод `.find()`, щоб знайти позицію кожної літери вашого повідомлення в алфавіті. Метод схожий на функцію, але належить об’єкту.

```py
sentence = 'My brain hurts!'
sentence.find('r')
```

Вгорі метод `.find()` *викликається на* `sentence` (рядок, у якому шукають), і `'r'` (символ, який потрібно знайти) передається як аргумент. Виклик `sentence.find('r')` поверне `4`, що є індексом першого входження `'r'` у `sentence`.

Наприкінці вашого коду викличте `.find()` на `alphabet` і передайте `'z'` як аргумент методу.

# --hints--

Ви повинні викликати метод `.find()`.

```js
assert.match(code, /find\s*\(.*\)/)
```

Ви повинні викликати метод `.find()` на змінній `alphabet`.

```js
assert.match(code, /alphabet\.find\s*\(.*\)/)
```


Ви повинні викликати метод `.find()` на змінній `alphabet` і передати `'z'` методу. Зверніть увагу, що виклик методу має бути на початку рядка.

```js
assert.match(code, /^alphabet\.find\s*\(\s*('|")z\1\s*\)/m)
```


Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
```
