---
id: 65520e6f2b9678799977f24d
title: Крок 18
challengeType: 20
dashedName: step-18
---

# --description--

Видаліть останній виклик `print()`. Потім, замість `text[0]`, передайте `text[0].lower()` як аргумент до вашого виклику `.find()` і подивіться на вихідні дані.

# --hints--

У вашому коді не повинно бути `print(text.lower())`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*text\.lower\s*\(\s*\)\s*\)/);
```

Ви все ще повинні вивести вашу змінну `index`.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m);
```

Вам слід оновити виклик `alphabet.find(text[0])`, щоб використовувати `text[0].lower()` як аргумент. Зверніть увагу, що виклик методу має бути на початку лінії.

```js
assert.match(code, /^index\s*=\s*alphabet\.find\s*\(\s*text\s*\[\s*\s*0\s*\]\s*\.lower\s*\(\s*\)\s*\)\s*(#.*)?$/m);
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
print(index)
print(text.lower())
--fcc-editable-region--
```
