---
id: 6553a755879b131a445e664c
title: Крок 44
challengeType: 20
dashedName: step-44
---

# --description--

Спробуйте призначити рядок `'Hello Zaira'` вашій змінній `text` і подивіться, що станеться в терміналі.

Ви побачите виняток `string index out of range`. Не хвилюйтеся, ви скоро зрозумієте, як це виправити!

# --hints--

У вас має бути змінна `text`.

```js
assert.match(code, /^text\s*=/m)
```

Ви маєте призначити рядок `'Hello Zaira'` вашій змінній `text`.

```js
assert.match(code, /^text\s*=\s*("|')Hello\sZaira\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
