---
id: 6553efd6ada3f42aa2d75448
title: Крок 53
challengeType: 20
dashedName: step-53
---

# --description--

Тепер ви знову повинні побачити вихідні дані. Хоч цей підхід і працює, він не значно покращує повторне використання коду. Повторні виклики вашої функції призведуть до однакового результату. Однак функції можна оголосити з *параметрами*, щоб надати їм гнучкості та можливість налаштування:

```py
def function_name(param_1, param_2):
    <code>
```

Параметри — це змінні, які ви можете використовувати всередині своєї функції. Функція може бути оголошена з різною кількістю параметрів. У наведеному вище прикладі `param_1` і `param_2` — це параметри.

Змініть оголошення вашої функції так, щоб вона приймала два параметри з іменами `message` і `offset`.

Після цього в терміналі з’явиться помилка. Ви побачите, як її виправити у наступних кроках.

# --hints--

Ваша функція `caesar` має приймати `message` і `offset` як параметри. Не забудьте розділити параметри комою.

```js
assert.match(code, /^def\s+caesar\s*\(\s*message\s*,\s*offset\s*\)\s*:/m)
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3
--fcc-editable-region--
def caesar():
--fcc-editable-region--
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

caesar()
```
