---
id: 6553f4f66099802c6ae94613
title: Крок 55
challengeType: 20
dashedName: step-55
---

# --description--

Наразі ваш код викликає `TypeError`, тому що функція `caesar` визначена з двома параметрами (`message` і `offset`), отже вона очікує виклик з двома *аргументами*.

Виклик `caesar()` без потрібних аргументів зупиняє виконання коду.

Надайте `message` і `offset` значення, передаючи `text` і `shift` як аргументи у виклик функції `caesar`.

# --hints--

Ви повинні передати `text` і `shift` як аргументи у виклик функції, включивши їх у круглі дужки. Не забудьте розділити аргументи комою.

```js
({
    test: () => assert(runPython(
        `_Node(_code).find_calls('caesar')[0].is_equivalent('caesar(text, shift)')`
    ))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira'
shift = 3

def caesar(message, offset):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
caesar()
--fcc-editable-region--
```
