---
id: 6553a44b1801991847d8cc69
title: Крок 47
challengeType: 20
dashedName: step-47
---

# --description--

Далі змініть свій виклик `print()`, щоб вивести `'encrypted text:', encrypted_text` і розмістіть його поза межами циклу `for`, щоб зашифрований рядок виводився один раз.

# --hints--

Ви повинні вивести `'encrypted text:', encrypted_text` після свого циклу for.

```js
({ test: () => assert.match(code, /^print\s*\(\s*("|')encrypted\stext:\1\s*,\s*encrypted_text\s*\)/m) })
```

У вашому коді не повинно бути `print('char:', char, 'encrypted text:', encrypted_text)`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.notMatch(commentless_code, /print\s*\(\s*("|')char:\1\s*,\s*char\s*,\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/)
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
