---
id: 6553a572f7a65718f1e42e18
title: Крок 49
challengeType: 20
dashedName: step-49
---

# --description--

Функція — це, по суті, багаторазовий блок коду. Ви вже зустрічали деякі вбудовані функції, як-от `print()`, `find()` і `len()`. Але ви також можете оголосити кастомні функції ось так:

```py
def function_name():
    <code>
```

Оголошення функції починається з ключового слова `def`, за яким іде ім’я функції — дійсне ім’я змінної — і пара круглих дужок. Оголошення завершується двокрапкою.

Одразу після вашої змінної `shift` оголосіть функцію з ім’ям `caesar` і відступіть усі наступні рядки, щоб надати вашій новій функції тіло.

# --hints--

Ви повинні використати ключове слово `def` для оголошення нової функції.

```js
assert.match(code, /^def\s+/m)
```

Ви повинні написати `caesar` як ім’я функції після ключового слова `def`. Не забудьте додати пробіл після `def`.

```js
assert.match(code, /^def\s+caesar/m)
```

Ви повинні додати пару круглих дужок після імені функції. Не забудьте фінальну двокрапку.

```js
assert.match(code, /^def\s+caesar\s*\(\s*\)\s*:/m)
```

Ви повинні відступити всі рядки після `shift = 3`, щоб вони стали тілом вашої нової функції.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {def} = __helpers.python.getDef(commentless_code, "caesar");    
    const replacement = def.replace(/print\s*\(\s*("|')plain\stext:\1\s*,\s*text\s*\)\s*print\s*\(\s*("|')encrypted\stext:\2\s*,\s*encrypted_text\s*\)/, "return encrypted_text")
    const py_code = `
text = "Hello Zaira"
shift = 3
${replacement}
caesar()
`
    const out = runPython(py_code, {});
    assert.equal(out, "khoor cdlud");
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3

alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = (index + shift) % len(alphabet)
        encrypted_text += alphabet[new_index]
print('plain text:', text)
print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
