---
id: 6553a7d8d05cbb1ae335a665
title: Крок 45
challengeType: 20
dashedName: step-45
---

# --description--

Коли цикл доходить до літери `Z`, сума `index + shift` перевищує останній індекс рядка `alphabet`. Тому `alphabet[new_index]` намагається використати недійсний індекс, що спричиняє викидання `IndexError`.

Ви можете помітити, що вихідні дані в терміналі зупиняються на пробілі безпосередньо перед `Z`, останнім `print` перед тим, як виникає помилка.

У цьому випадку оператор за модулем (`%`) можна використати, щоб повернути остачу від ділення двох чисел. Наприклад: `5 % 2` дорівнює `1`, тому що 5, поділене на 2, має частку 2 і остачу 1.

Оберніть `index + shift` у круглі дужки, і застосуйте оператор за модулем до виразу з `26`, що є довжиною `alphabet`.

# --hints--

У вашій інструкції `new_index = (index + shift) % 26` має бути `else`.

```js
const commentless_code = __helpers.python.removeComments(code);  
const {block_body} = __helpers.python.getBlock(commentless_code, /else/);
assert(block_body.match(/new_index\s*=\s*\(\s*index\s*\+\s*shift\s*\)\s*%\s*26/));
```

Ваш код містить недійсний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    if char == ' ':
        encrypted_text += char
    else:
        index = alphabet.find(char)
        new_index = index + shift
        encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
