---
id: 65549f90cf78131c96ebcf28
title: Крок 61
challengeType: 20
dashedName: step-61
---

# --description--

При написанні коду важлива читабельність. Коментарі служать ефективними нотатками, які пояснюють логіку вашого коду. Вони стають у пригоді, коли ви повертаєтеся до проєкту через деякий час, а також допомагають вашим колегам зрозуміти код.

У Python ви можете написати коментар, використовуючи `#`. Все, що йде після `#`, не буде виконано.

Перед вашою інструкцією `if` додайте коментар із текстом `Append space to the message`.

# --hints--

Ви повинні додати коментар із текстом `Append space to the message`. Не забудьте символ `#` на початку.

```js
({ test: () => assert.match(code, /#\s*Append\sspace\sto\sthe\smessage/) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        
        if char == ' ':
            encrypted_text += char
--fcc-editable-region--            
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
