---
id: 65549561463f0016876e852c
title: Крок 60
challengeType: 20
dashedName: step-60
---

# --description--

Оскільки ваш ключ коротший за текст, який потрібно буде зашифрувати, вам доведеться повторювати його, щоб отримати цілий зашифрований текст. На початку тіла вашої функції оголосіть змінну `key_index` і призначте їй значення нуль.

# --hints--

Ви повинні оголосити змінну з назвою `key_index` на початку тіла вашої функції.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");
assert(function_body.match(/^\s*key_index\s*=/));
```

Ви повинні призначити `0` вашій змінній `key_index`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {function_body} = __helpers.python.getDef(commentless_code, "vigenere");   
assert(function_body.match(/key_index\s*=\s*0/));
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)
--fcc-editable-region--
```
