---
id: 6554a49a4f782f208abcc87e
title: Крок 63
challengeType: 20
dashedName: step-63
---

# --description--

Вам потрібно збільшити лічильник `key_index` для наступної ітерації.  
Для цього після лінії, яку ви щойно додали, і в тому самому блоці коду використайте оператор додавання з присвоєнням, щоб збільшити `key_index` на одиницю.

# --hints--

Ви повинні використати оператор `+=`, щоб додати `1` до `key_index` всередині тіла умови `else`.

```js
const commentless_code = __helpers.python.removeComments(code);
const {block_body} = __helpers.python.getBlock(commentless_code, "else");
assert(block_body.match(/key_char\s*=\s*key\s*\[\s*key_index\s*%\s*len\s*\(\s*key\s*\)\s*\]\s*key_index\s*\+=\s*1/));
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
--fcc-editable-region--
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:
            key_char = key[key_index % len(key)]
--fcc-editable-region--
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', message)
    print('encrypted text:', encrypted_text)

```
