---
id: 6554d25dc5ceaa354307a77e
title: Крок 70
challengeType: 20
dashedName: step-70
---

# --description--

Шифрування та розшифрування — це протилежні процеси, і ваша функція може виконувати обидва з невеликими змінами.

Додайте третій параметр з назвою `direction` до визначення вашої функції. Також перетворіть на коментарі останні дві лінії коду, щоб уникнути помилок у консолі.

# --hints--

Ви повинні перетворити останні дві лінії у вашому коді на коментарі. Додайте `#` на початок кожної лінії.

```js
({ test: () => assert.match(code, /#\s*encryption\s*=\s*vigenere\s*\(\s*text\s*,\s*custom_key\s*\)\s*#\s*print\s*\(\s*encryption\s*\)/) })
```

Ви повинні додати `direction` як третій параметр вашої функції.

```js
({ test: () => assert(runPython(`
    import inspect    
    sig = str(inspect.signature(vigenere))
    sig == '(message, key, direction)'
  `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello Zaira'
custom_key = 'python'

def vigenere(message, key):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in message.lower():
    
        # Append space to the message
        if char == ' ':
            encrypted_text += char
        else:        
            # Find the right key character to encode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset) % len(alphabet)
            encrypted_text += alphabet[new_index]
    
    return encrypted_text
    
encryption = vigenere(text, custom_key)
print(encryption)
--fcc-editable-region--
```
