---
id: 6555e1bfcd374e18c6be8e58
title: Крок 91
challengeType: 20
dashedName: step-91
---

# --description--

У Python є спосіб легко форматувати рядки. *f-рядки* дозволяють вам інтерполювати значення у своїх рядках.

Інтерполяція означає написання заповнювачів, які будуть замінені на вказані значення під час запуску програми. Наприклад, ви можете отримати той самий результат `'Encrypted text: ' + text` за допомогою `f'Encrypted text: {text}'`. Для створення f-рядка потрібно поставити `f` перед лапками і написати змінні або вирази, які хочете інтерполювати, у фігурних дужках.

Змініть перший виклик `print()`, щоб використати f-рядок.

# --hints--

Вам слід змінити виклик `print('Encrypted text: ' + text)`, щоб він виводив той самий результат, використовуючи f-рядок.

```js
({ test: () => assert.match(code, /print\s*\(\s*f("|')Encrypted\stext:\s\{\s*text\s*\}\1\s*\)/) })
```

# --seed--

## --seed-contents--

```py
text = 'mrttaqrhknsw ih puggrur'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message

def encrypt(message, key):
    return vigenere(message, key)
    
def decrypt(message, key):
    return vigenere(message, key, -1)
--fcc-editable-region--
print('Encrypted text: ' + text)
print('Key: ' + custom_key)
--fcc-editable-region--
#decryption = decrypt(encryption, custom_key)
#print(decryption)
```
