---
id: 6555231eac4d9f19bd3d44b5
title: Крок 83
challengeType: 20
dashedName: step-83
---

# --description--

Ключове слово `pass` можна використовувати як заповнювач для майбутнього коду. Воно не має жодного впливу на ваш код, але може врятувати вас від помилок у разі неповного коду:

```py
def foo():
    pass
```

Виклик `vigenere` з `1` для шифрування та `-1` для дешифрування цілком прийнятний, але може бути трохи незрозумілим. Створіть нову функцію з назвою `encrypt`, яка прийматиме параметри `message` і `key`, і використайте `pass` для заповнення тіла функції.


# --hints--

Ви повинні оголосити нову функцію з назвою `encrypt`. Пам’ятайте, що всередині тіла потрібно використовувати ключове слово `pass`.

```js
({ test: () => assert(runPython(`
    import inspect
    inspect.isfunction(encrypt)
  `))
})
```

Ваша функція `encrypt` має приймати `message` і `key` як параметри. Порядок має значення.

```js
({ test: () => assert(runPython(`
    import inspect
    sig = inspect.signature(encrypt)
    str(sig) == "(message, key)"
  `))
})
```

# --seed--

## --seed-contents--

```py
text = 'Hello Zaira!'
custom_key = 'python'

def vigenere(message, key, direction=1):
    key_index = 0
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    final_message = ''

    for char in message.lower():

        # Append any non-letter character to the message
        if not char.isalpha():
            final_message += char
        else:        
            # Find the right key character to encode/decode
            key_char = key[key_index % len(key)]
            key_index += 1

            # Define the offset and the encrypted/decrypted letter
            offset = alphabet.index(key_char)
            index = alphabet.find(char)
            new_index = (index + offset*direction) % len(alphabet)
            final_message += alphabet[new_index]
    
    return final_message
--fcc-editable-region--

encryption = vigenere(text, custom_key)
print(encryption)
decryption = vigenere(encryption, custom_key, -1)
print(decryption)
--fcc-editable-region--
```
