---
id: 655a2a7210094920069b117c
title: Крок 40
challengeType: 20
dashedName: step-40
---

# --description--

Оператори порівняння дозволяють порівнювати два об’єкти на основі їхніх значень. Ви можете використати оператор порівняння, розмістивши його між об’єктами, які хочете порівняти.
Вони повертають *булеве* значення — а саме `True` або `False` — залежно від істинності виразу.

У Python є такі оператори порівняння:

<table>
  <thead>
    <tr>
      <th>Оператор</th>
      <th>Значення</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td>==</td>
      <td>Рівний</td>
    </tr>
    <tr>
      <td>!=</td>
      <td>Не рівний</td>
    </tr>
    <tr>
      <td>&gt;</td>
      <td>Більший за</td>
    </tr>
    <tr>
      <td>&lt;</td>
      <td>Менший за</td>
    </tr>
    <tr>
      <td>&gt;=</td>
      <td>Більший або рівний</td>
    </tr>
    <tr>
      <td>&lt;=</td>
      <td>Менший або рівний</td>
    </tr>
  </tbody>
</table>

На початку тіла вашого циклу виведіть результат порівняння `char` з пробілом (`' '`). Для цього використайте оператор рівності `==`.

# --hints--

Ви повинні порівняти `char` з пробілом, використовуючи оператор рівності всередині вашого циклу `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/char\s*==\s*("|')\s\1/));
  }
})
```

Ви повинні вивести результат порівняння `char` з пробілом всередині вашого циклу `for`.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

Ви повинні вивести результат порівняння `char` з пробілом на початку вашого циклу.

```js
({ test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    const {block_body} = __helpers.python.getBlock(commentless_code, /for\s+char\s+in\s+text\.lower\s*\(\s*\)\s*/);
    assert(block_body.match(/^\s+print\s*\(\s*char\s*==\s*("|')\s\1\s*\)/));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
encrypted_text = ''

for char in text.lower():
    index = alphabet.find(char)
    new_index = index + shift
    encrypted_text += alphabet[new_index]
    print('char:', char, 'encrypted text:', encrypted_text)
--fcc-editable-region--
```
