---
id: 65687d2f8c7ee27b0446cef3
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

Код, який виконується на кожній ітерації — розміщений після `:` — становить тіло циклу. Цей код має бути відступлений. У Python рекомендується використовувати 4 пробіли на рівень відступу. Цей відступлений рівень є блоком коду.

```py
for i in text:
    print(i)
```

Python покладається на відступи, щоб позначити блоки коду. Двокрапка в кінці рядка сигналізує, що далі йтиме новий відступлений блок коду.

Отже, коли після фінальної двокрапки не знаходиться відступленого блоку, виконання коду зупиняється і виникає `IndentationError`. Цей код не покаже вихідні дані, а замість цього викличе `IndentationError`:

```py
for i in text:
print(i)
```

Додайте тіло вашому циклу `for`, додавши виклик `print(i)`. Не забудьте відступити тіло циклу.

# --hints--

Ви маєте додати `print(i)` до тіла вашого циклу `for`. Зверніть увагу на відступ.

```js
assert.match(code, /^for\s+i\s+in\s+text:\s+print\s*\(\s*i\s*\)/m)
```

Ваш код містить неправильний синтаксис і/або неправильний відступ.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for i in text:

--fcc-editable-region--
```
