---
id: 65688b5a1655a7a6caede847
title: Крок 31
challengeType: 20
dashedName: step-31
---

# --description--

Коли ви намагаєтеся змінити окремі символи рядка, як робили на попередньому кроці, ви отримуєте `TypeError`, що виникає, коли в коді використовується об’єкт невідповідного типу.

Як видно з повідомлення про помилку, рядки не підтримують присвоєння елементів, бо вони незмінні. Проте змінній можна присвоїти інший рядок:

```py
message = 'Hello World'
message = 'Hello there!'
```

Видаліть лінію `text[0]` і присвойте змінній `text` рядок `'Albatross'`.

# --hints--

У вашому коді не повинно бути `text[0] = 'C'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.isFalse( /text\s*\[\s*0\s*\]\s*\=\s*("|')\w\1/.test(commentless_code));
```

Ви маєте присвоїти змінній `text` рядок `'Albatross'`. Не змінюйте `text = 'Hello World'`.

```js
({test: () => {
    const commentless_code = __helpers.python.removeComments(code);
    assert.match(commentless_code.replace(/\r/g, ''), /text\s*=\s*("|')Hello World\1\s+text\s*=\s*("|')Albatross\2/m);
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
text[0] = 'C'
--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
