---
id: 65688a50e6c998a21d8e41d3
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Рядки є незмінними, тобто їх не можна змінити після створення. Наприклад, ви могли подумати, що наступний код змінює значення `my_string` на рядок `'train'`, але це не є дійсним:

```py
my_string = 'brain'
my_string[0] = 't'
```

Підтвердіть це, використовуючи дужкову нотацію для доступу до першої літери в `text` і спробуйте змінити її на символ на свій вибір. Ви побачите, що вивід зникне, а з’явиться помилка.

# --hints--

У вашому коді має залишатися `text = 'Hello World'`.

```js
assert.match(code, /^text\s*=\s*("|')Hello World\1/m)
```

Ви повинні отримати доступ до першої літери в `text` за допомогою `text[0]`.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*/m)
```

Використайте оператор `=`, щоб призначити символ на свій вибір `text[0]`. Не забудьте взяти символ у одинарні або подвійні лапки.

```js
assert.match(code, /^text\s*\[\s*0\s*\]\s*\=\s*("|').\1/m)
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'

for char in text.lower():
    index = alphabet.find(char)
    print(char, index)
    new_index = index + shift

```
