---
id: 656886ec4e4d9b96ded3c565
title: Крок 25
challengeType: 20
dashedName: step-25
---

# --description--

Змінна ітерації може мати будь-яке дійсне ім’я, але зручно давати їй осмислену назву.

Перейменуйте вашу змінну `i` на `char`.

# --hints--

Ви повинні замінити `i` на `char`. 

```js
assert.match(code, /^for\s+char\s+in\s+text:\s+print\s*\(\s*char\s*\)/m)
```

У вашому коді є недійсний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
--fcc-editable-region--
for i in text:
    print(i)
--fcc-editable-region--
```
