---
id: 65952b70924edb449faf1950
title: Крок 19
challengeType: 20
dashedName: step-19
---

# --description--

Оголосіть нову змінну з іменем `shifted`. Використайте дужкову нотацію, щоб отримати доступ до значення `alphabet` за індексом `index` і призначте його своїй новій змінній.

# --hints--

Ви повинні оголосити змінну з іменем `shifted`. Зверніть увагу, що ім'я змінної має бути на початку лінії.

```js
assert.match(code, /^shifted\s*=/m)
```

Ви повинні призначити `alphabet[index]` своїй змінній `shifted`.

```js
assert.match(code, /^shifted\s*=\s*alphabet\s*\[\s*index\s*\]\s*(#.*)?$/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
--fcc-editable-region--
```
