---
id: 659530ca823d6c47f61daa8a
title: Крок 20
challengeType: 20
dashedName: step-20
---

# --description--

Виведіть вашу змінну `shifted`.

# --hints--

Ви маєте вивести вашу змінну `shifted`. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*shifted\s*\)/m)
```

У вашому коді є помилка синтаксису та/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index]
--fcc-editable-region--
```
