---
id: 65951b637ddec63611c482c0
title: Крок 3
challengeType: 20
dashedName: step-3
---

# --description--

Ви можете використовувати вбудовану функцію `print()`, щоб вивести вихідні дані вашого коду в термінал.

Функції — це повторно використовувані блоки коду, які ви можете викликати, щоб запустити їхній код, коли це потрібно. Щоб викликати функцію, достатньо написати пару круглих дужок поруч із її назвою. Ви дуже скоро дізнаєтеся більше про функції.

Поки що перейдіть на новий рядок і додайте порожній виклик функції `print()`. Ви поки що не повинні бачити жодних вихідних даних.

# --hints--

У вашому коді має бути `print()`. Зверніть увагу, щоб розмістити виклик функції на початку рядка.

```js
assert.match(code, /^print\s*\(\s*\)/m)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'

--fcc-editable-region--
```
