---
id: 659526f3e5424b3eac4b6aa7
title: Крок 16
challengeType: 20
dashedName: step-16
---

# --description--

Далі виведіть змінну `index` у консоль.

# --hints--

Ви маєте вивести свою змінну `index`. Зверніть увагу, щоб виклик функції був на початку лінії.

```js
assert.match(code, /^print\s*\(\s*index\s*\)/m)
```

У вашому коді є неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0])
--fcc-editable-region--
```
