---
id: 659db3e3670d3e712be82593
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Повторювати процес пошуку літери в алфавіті та визначення зсунутої літери для кожного символу в `text` може бути втомливо. На щастя, ви можете спростити це, використавши цикл.

Поки що видаліть усі рядки коду нижче оголошення змінної `alphabet`.

# --hints--

У вашому коді має залишитися `text = 'Hello World'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^text\s*=\s*("|')Hello World\1/m)
```

У вашому коді має залишитися `shift = 3`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^shift\s*=\s*3/m)
```

У вашому коді має залишитися `alphabet = 'abcdefghijklmnopqrstuvwxyz'`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.match(commentless_code, /^alphabet\s*=\s*("|')abcdefghijklmnopqrstuvwxyz\1/m)
```

Ви маєте видалити змінну `index` та її значення.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /index\s*=/)
```

У вашому коді не повинно бути `print(index)`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*index\s*\)/)
```

Ви маєте видалити змінну `shifted` та її значення.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /shifted\s*=/)
```

У вашому коді не повинно бути `print(shifted)`.

```js
const commentless_code = __helpers.python.removeComments(code);
assert.notMatch(commentless_code, /print\s*\(\s*shifted\s*\)/)
```

Ваш код містить неправильний синтаксис і/або неправильне відступлення.

```js
({test: () => assert(true) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
index = alphabet.find(text[0].lower())
print(index)
shifted = alphabet[index + shift]
print(shifted)
--fcc-editable-region--
```
