---
id: 6624e0ec0bf47017eed84b1f
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Перший тип шифру, який ви збираєтеся створити, називається *шифр Цезаря*. Конкретно, ви візьмете кожну літеру у вашому повідомленні, знайдете її позицію в алфавіті, візьмете літеру, що знаходиться через 3 позиції в алфавіті, і заміните оригінальну літеру на нову.

Щоб реалізувати це, ви використаєте метод `.find()`, про який йшлося на попередньому кроці. Змініть свій існуючий виклик `.find()`, передаючи йому `text[0]` як аргумент замість `'z'`.

# --hints--

Ви повинні змінити свій існуючий виклик `alphabet.find('z')`, передаючи методу `text[0]`.

```js
assert.match(code, /^alphabet\.find\s*\(\s*text\s*\[\s*0\s*\]\s*\)/m)
``` 

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
text = 'Hello World'
shift = 3
alphabet = 'abcdefghijklmnopqrstuvwxyz'
alphabet.find('z')
--fcc-editable-region--
```
