---
id: 66605a64a9822f01ec5ae8ce
title: Крок 50
challengeType: 20
dashedName: step-50
---

# --description--

В Python *область видимості* змінної визначає, де можна отримати доступ до цієї змінної:

- Змінні, визначені поза межами функції, мають *глобальну* область видимості і до них можна звертатися з будь-якої частини вашого коду.

- Змінні, визначені всередині функції, є *локальними* для цієї функції і до них не можна звертатися поза межами функції.

Щоб побачити це на практиці, спробуйте вивести змінну `alphabet` наприкінці вашого коду. Це викличе виняток `NameError`.

Ви повинні побачити повідомлення про помилку, що `alphabet` не визначено. Це тому, що `alphabet` визначена всередині функції `caesar` і недоступна поза межами цієї функції.

# --hints--

Ви повинні спробувати вивести змінну `alphabet` поза межами функції caesar.

```js
assert.match(code, /^print\(\s*alphabet\s*\)/m)
```

# --seed--

## --seed-contents--

```py

text = 'Hello Zaira'
shift = 3

def caesar():
    alphabet = 'abcdefghijklmnopqrstuvwxyz'
    encrypted_text = ''

    for char in text.lower():
        if char == ' ':
            encrypted_text += char
        else:
            index = alphabet.find(char)
            new_index = (index + shift) % len(alphabet)
            encrypted_text += alphabet[new_index]
    print('plain text:', text)
    print('encrypted text:', encrypted_text)

--fcc-editable-region--

--fcc-editable-region--
```
