---
id: 69899123889dcf5b19be1847
title: Implementieren Sie den Algorithmus der Tiefensuche
challengeType: 26
dashedName: implement-depth-first-search-algorithm
---

# --description--

In diesem Labor implementieren Sie einen Graphdurchlauf-Algorithmus namens <dfn>Tiefensuche</dfn> (DFS).

Während eine <dfn>Breitensuche</dfn> (BFS) Nachbarn Ebene für Ebene erkundet, taucht die <dfn>Tiefensuche</dfn> so tief wie möglich entlang eines einzelnen Pfads von Kanten ein, bevor sie zurückverfolgt.

Sobald der Algorithmus das Ende eines Pfads erreicht (ein Knoten ohne unbesuchte Nachbarn), kehrt er zum zuletzt besuchten Knoten zurück, der noch unerforschte Kanten hat, und setzt die Suche von dort fort.

Um diesen Algorithmus iterativ zu implementieren, sollten Sie einen **Stack** verwenden. In JavaScript eignet sich ein Standard-`Array` perfekt als Stack, indem die `.push()`- und `.pop()`-Methoden genutzt werden, gemäß dem **Last-In-First-Out (LIFO)**-Prinzip. Das stellt sicher, dass der zuletzt zum Stack hinzugefügte Nachbar als nächstes erkundet wird.

Eine einfache Ausgabe dieses Algorithmus ist eine Liste von Knoten, die von einem gegebenen Knoten aus erreichbar sind. Daher sollten Sie auch die besuchten Knoten im Auge behalten.

**Ziel**: Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `dfs` definieren.
1. Die `dfs`-Funktion sollte zwei Argumente annehmen:

    - `graph`: Eine Adjazenzmatrix (ein Array von Arrays).
    - `root`: Eine numerische Knotenbezeichnung (der Startindex), der numerische Wert des Knotens zwischen `0` und `n - 1`, wobei `n` die Gesamtanzahl der Knoten im Graphen ist.

1. Die Funktion sollte ein Array zurückgeben, das alle von dem Startknoten aus erreichbaren Knotenbezeichnungen enthält.

# --hints--

Sie sollten eine Funktion mit dem Namen `dfs` haben, die zwei Argumente annimmt.

```js
assert.isFunction(dfs);
assert.lengthOf(dfs, 2);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1)` sollte eine Liste mit `1`, `2`, `3` und `0` zurückgeben.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 1), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3)` sollte eine Liste mit `1`, `2`, `3` und `0` zurückgeben.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 1], [0, 0, 1, 0]], 3), [1, 2, 3, 0]);
```

`dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3)` sollte `[3]` zurückgeben.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 1, 0], [0, 1, 0, 0], [0, 0, 0, 0]], 3), [3]);
```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3)` sollte eine Liste mit `3` und `2` zurückgeben.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 3), [3, 2]);

```

`dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0)` sollte eine Liste mit `0` und `1` zurückgeben.

```js
assert.sameMembers(dfs([[0, 1, 0, 0], [1, 0, 0, 0], [0, 0, 0, 1], [0, 0, 1, 0]], 0), [0, 1]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dfs(graph, root) {
  let stack = [root];
  let visited = [];

  while (stack.length > 0) {
    let current = stack.pop();

    if (!visited.includes(current)) {
      visited.push(current);
      let neighbors = graph[current];
      
      for (let i = 0; i < neighbors.length; i++) {
        if (neighbors[i] === 1) {
          stack.push(i);
        }
      }
    }
  }
  return visited;
}
```
