---
id: 686d749d6f32f708fdf2b7ad
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在之前的講座中，你已經瞭解了 `svg` 元素及其在項目中常被用作圖標的用法。 在實際代碼庫中，通常會依賴圖標庫，因此無需從頭創建 `svg` 元素。 然而，在本工作坊中，你將通過構建一個心形圖標來學習 `svg` 元素內部的核心屬性。

首先在頁面上創建一個 `svg` 元素。

# --hints--

頁面上應包含一個 `svg` 元素。

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 

    --fcc-editable-region--
  </body>
</html>
```
