---
id: 686d749d6f32f708fdf2b7ad
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在之前的課程中，你學習了 `svg` 元素以及它們如何經常用作專案中的圖示。在真實的程式碼庫中，你通常會依賴圖示庫，因此不需要從頭創建 `svg` 元素。然而，在這個工作坊中，你將建造一個愛心圖示，以了解用於 `svg` 元素內的核心屬性。

首先在頁面上建立一個 `svg` 元素。

# --hints--

你應該在頁面上有一個 `svg` 元素。

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
