---
id: 686daa7ed79ceacd0b264e7d
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

你應該在你的 `svg` 元素內巢狀一個 `path` 元素，以賦予影像形狀。

創建一個 `path` 元素。

# --hints--

你應該在你的 `svg` 元素內巢狀一個 `path` 元素。

```js
const path = document.querySelector('svg path');
assert.exists(path);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    <svg>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </svg>
  </body>
</html>
```
