---
id: 686daa7ed79ceacd0b264e80
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

你越來越接近了，現在看看這個範例：

```html
<svg viewBox="0 0 50 50">
</svg>
```

`viewBox` 屬性控制 SVG 中可見的影像部分。

- 前兩個數字（`0 0`）設定了 `viewBox` 的起始位置——左上角（x 和 y）。
- 接下來的兩個數字（`50 50`）定義了 `viewBox` 的寬度和高度。

將 `viewBox` 屬性設為 `0 0 24 24`。

# --hints--

你應該有一個 `viewBox` 屬性。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

你應該將 `viewBox` 的 x 位置設為 0。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

你應該將 `viewBox` 的 y 位置設為 0。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

你應該將 `viewBox` 寬度設為 24。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

你應該將 `viewBox` 高度設為 24。

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
