---
id: 686daa7ed79ceacd0b264e80
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

`path` 元素需要定義其形狀。 這就該 `d` 屬性發揮作用了。 它用於生成一系列命令字母和數字，從而繪製出形狀。

這些字母代表諸如移動至、繪製直線和閉合等命令，而數字則代表座標。

將心形圖標的 `d` 屬性設置爲 `M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z`。

# --hints--

你的 `path` 元素應該有一個 `d` 屬性。

```js
const path = document.querySelector('path');
const d = path.getAttribute('d');
assert.exists(d);
```

你應該將 `d` 屬性設置爲 `M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z`

```js
const path = document.querySelector('path');
const d = path.getAttribute('d');
assert.strictEqual(
  d,
  'M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z'
);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24" viewBox="0 0 24 24">
      <path></path>
    </svg>
    --fcc-editable-region--
  </body>
</html>
```
