---
id: 66629f407d679d3105e8317f
title: Costruisci un gioco carta forbice sasso
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
Il tuo gioco sarà giocato contro il computer. Scriverai una funzione che restituisce casualmente `"rock"`, `"paper"` o `"scissors"`.

Non devi preoccuparti della parte front-end del gioco. Scriverai solo la logica del gioco. Apri il `script.js` e inizia a seguire le istruzioni.

**Storie utente:**

1. Dovresti avere una funzione chiamata `getComputerChoice`.

1. La tua funzione `getComputerChoice` dovrebbe restituire `"rock"`, `"paper"` o `"scissors"` in modo casuale.

**Suggerimento:** Il metodo [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) restituisce un numero casuale maggiore o uguale a 0 e minore di 1. Pensa a come puoi usarlo per restituire condizionalmente una delle scelte multiple.

Il tuo gioco sarà giocato da un giocatore umano. Scriverai una funzione che prende la scelta dell'utente e la restituisce.

1. Crea una funzione chiamata `getHumanChoice`.

1. Scrivi il codice in modo che `getHumanChoice` restituisca una delle scelte valide a seconda di ciò che l'utente inserisce.

**Suggerimento:** Usa il metodo [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) per ottenere l'input dell’utente.

Il tuo gioco terrà traccia del punteggio del giocatore. Scriverai variabili per tenere traccia del punteggio del giocatore.

1. Crea due nuove variabili chiamate `humanScore` e `computerScore` nell’ambito globale.

1. Inizializza queste variabili con il valore `0`.

Il tuo gioco sarà giocato round per round. Scriverai una funzione che prende come argomenti le scelte del giocatore umano e del computer, gioca un singolo round, incrementa il punteggio del vincitore del round e registra un annuncio del vincitore.

1. Crea una nuova funzione chiamata `playRound`.

1. Definisci due parametri per `playRound`. Il primo parametro `humanChoice` e il secondo parametro `computerChoice`. Usa questi due parametri per prendere come argomenti le scelte dell’umano e del computer.

1. Rendi il parametro `humanChoice` della tua funzione insensibile alle maiuscole in modo che i giocatori possano inserire `"rock"`, `"ROCK"`, `"RocK"` o altre variazioni.

1. Scrivi il codice per la tua funzione `playRound` che restituisce un valore stringa che rappresenta il vincitore del round.

- Se c’è un pareggio, dovrebbe restituire `"It's a tie!"`.

- Se il giocatore vince, dovrebbe restituire `"You win! [player choice] beats [computer choice]"`.

- Se il computer vince, dovrebbe restituire `"You lose! [computer choice] beats [player choice]"`.

1. Incrementa la variabile `humanScore` o `computerScore` in base al vincitore del round.

Il tuo gioco giocherà 3 round. Scriverai una funzione chiamata `playGame` che chiama `playRound` per giocare 3 round, tiene traccia dei punteggi e dichiara un vincitore alla fine.

1. Crea una nuova funzione chiamata `playGame`.

1. Crea un loop che giochi 3 round e chiami la funzione `playRound` ogni volta con le funzioni delle scelte dell’umano e del computer come argomenti.

1. Alla fine del gioco, restituisci il vincitore del gioco in base a chi ha vinto più round.

- Se il giocatore umano vince più round del computer, restituisci un messaggio che dice `"You win the game!"`.

- Se il computer vince più round del giocatore umano, restituisci un messaggio che dice `"You lose the game!"`.

# --hints--

Dovresti avere una funzione chiamata `getComputerChoice`

```js
assert.isFunction(getComputerChoice);
```

La tua funzione `getComputerChoice` dovrebbe restituire `"rock"`, `"paper"` o `"scissors"` in modo casuale.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

Dovresti avere una funzione chiamata `getHumanChoice`

```js
assert.isFunction(getHumanChoice);
```

Dovresti avere due variabili chiamate `humanScore` e `computerScore` nell’ambito globale.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

Dovresti avere una funzione chiamata `playRound`

```js
assert.isFunction(playRound);
```

La tua funzione `playRound` dovrebbe prendere come argomenti le scelte del giocatore umano e del computer con i parametri `humanChoice` e `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

La tua funzione `playRound` dovrebbe essere insensibile alle maiuscole in modo che i giocatori possano inserire `"rock"`, `"ROCK"`, `"RocK"` o altre variazioni.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

La tua funzione `playRound` dovrebbe restituire la stringa `"It's a tie!"` se le scelte dell’umano e del computer sono uguali.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

La tua funzione `playRound` dovrebbe restituire la stringa `"You win! [player choice] beats [computer choice]"` se il giocatore vince.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

La tua funzione `playRound` dovrebbe restituire la stringa `"You lose! [computer choice] beats [player choice]"` se il computer vince.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

La tua funzione `playRound` dovrebbe incrementare la variabile `humanScore` o `computerScore` in base al vincitore del round.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

Dovresti avere una funzione chiamata `playGame`.

```js
assert.isFunction(playGame);
```

Dovresti usare un loop per giocare 3 round.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

Dovresti restituire il vincitore del gioco in base a chi ha vinto più round.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```
