---
id: 66731cd027ef3acb155669f5
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will learn how to work with basic HTML elements like headings, paragraphs, and lists by building a cat photo app.

The first element you will learn about is the `h1` element. The `h1` element is a heading element and is used for the main heading of a webpage.

```html
<h1>This is a main heading</h1>
```

Add an `h1` element with the text of `CatPhotoApp` and watch the change in the browser preview.

When you are done, press the "Check Your Code" button to see if it's correct.

# --hints--

代碼中的文本應該是 `CatPhotoApp`。 你可能想要檢查你的拼寫。

```js
assert.match(code, /catphotoapp/i);
```

你的 `h1` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該有一個結束標籤。 Closing tags have this syntax: `</elementName>`.

```js
assert.match(code, /<\/h1\>/);
```

你的 `h1` 元素的文本應爲 `CatPhotoApp`。 你可能忽略了此文本，或者有拼寫錯誤，或沒有把它放在 `h1` 元素的開始和結束標籤之間。

```js
assert.equal(document.querySelector('h1').innerText.toLowerCase(), 'catphotoapp');
```

你似乎使用了一個瀏覽器擴展，正在修改頁面。 一定要關閉所有瀏覽器擴展。

```js
assert.isAtMost(document.querySelectorAll('script').length, 2);
assert.equal(document.querySelectorAll('style').length, 0);
assert.equal(document.querySelectorAll('link').length, 0);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```
