---
id: 6690bebfd5c918179ea4fd3a
title: 步驟 9
challengeType: 0
dashedName: step-9
---

# --description--

All `img` elements should have an `alt` attribute. The `alt` attribute's text is used for screen readers to improve accessibility and is displayed if the image fails to load.

這是一個帶有 `alt` 屬性的 `img` 元素的示例：

```html
<img src="cat.jpg" alt="A cat">
```

Inside the `img` element, add an `alt` attribute with this text:

`A cute orange cat lying on its back`

# --hints--

你的代碼應該有一個 `img` 元素。 You removed the `img` element from an earlier step.

```js
assert.exists(document.querySelector('img'));
```

Your `img` element does not have an `alt` attribute. 請檢查在開始標籤的名稱後面要有一個空格，且/或所有的屬性名稱前面也要有一個空格。

```js
assert.isTrue(document.querySelector('img').hasAttribute('alt'));
```

Your `img` element's `alt` attribute value is set to something other than 'A cute orange cat lying on its back'. Make sure the `alt` attribute's value is surrounded with quotation marks.

```js
const altText = document
  .querySelector('img')
  .alt.toLowerCase()
  .replace(/\s+/g, ' ');
assert.match(altText, /A cute orange cat lying on its back\.?$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more cat photos in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg">
--fcc-editable-region--
    </main>
  </body>
</html>
```


