---
id: 6690be42942f421648b119eb
title: 步驟7
challengeType: 0
dashedName: step-7
---

# --description--

You can add images to your website by using the `img` element. `img` elements have an opening tag without a closing tag. An element without a closing tag is known as a <dfn>void element</dfn>.

在 `p` 元素的下方添加一個 `img` 元素。 此時，沒有圖片會出現在瀏覽器中。

# --hints--

你的 `img` 元素應該有一個開始標籤。 開始標籤的書寫語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('img'));
```

你的 `img` 元素應該沒有結束標籤 `</img>`。

```js
assert.notMatch(code, /<\/img\>/);
```

應該只有一個 `img` 元素。 刪除多餘的。

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

你的 `img` 元素應該在 `p` 元素的下方。 你把順序寫錯了。

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more cat photos in our gallery.</p>

--fcc-editable-region--
    </main>
  </body>
</html>
```


