---
id: 66984484deb84b3e21e5b3d7
title: 步驟-23
challengeType: 0
dashedName: step-23
---

# --description--

The `figure` element represents self-contained content and will allow you to associate an image with a caption.

將剛剛添加的圖像嵌套在 `figure` 元素中。

# --hints--

你的 `figure` 元素應該有一個開始標籤。 起始標籤的書寫語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('figure'));
```

你的 `figure` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert.match(code, /<\/figure\>/);
```

在第二個 `section` 元素的結束標籤上方應該有一個 `figure` 元素。

```js
assert.equal(document.querySelectorAll('section')[1].lastElementChild.nodeName, 'FIGURE');
```

Lasagna `img` 元素應該嵌套在 `figure` 元素中。

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--

        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```


