---
id: 66984e8ab54bee517207be42
title: 步驟 34
challengeType: 0
dashedName: step-34
---

# --description--

Nest a `p` element with the text `No Copyright - freeCodeCamp.org` within the `footer` element.

# --hints--

你要麼刪除了 `footer` 元素，要麼它缺少開始標籤或結束標籤。

```js
assert.exists(document.querySelector('footer'));
assert.match(code, /<\/footer>/);
```

你的 `footer` 元素應該有一個 `p` 元素。 確保爲 `p` 元素添加開始標籤和結束標籤。

```js
assert.exists(document.querySelector('footer > p'));
```

你的 `footer` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
const pElemClosingTags = code.match(/<\/p\>/g);
assert.exists(pElemClosingTags);
assert.lengthOf(pElemClosingTags, 2);
```

你的 `p` 元素的文本應該是 `No Copyright - freeCodeCamp.org`。 你可能忽略了文本，或者有拼寫錯誤，或沒有把它放在 `p` 元素的開始和結束標籤之間。

```js
const extraSpacesRemoved = document.querySelectorAll('footer > p')[0].innerText.replace(/\s+/g, ' ');
assert.match(extraSpacesRemoved, /No Copyright - freeCodeCamp\.org$/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
--fcc-editable-region--
    <footer>

    </footer>
--fcc-editable-region--
  </body>
</html>
```


