---
id: 6698417821794837f14ca275
title: 第 14 步
challengeType: 0
dashedName: step-14
---

# --description--

To open links in a new tab, you can use the `target` attribute on the anchor (`a`) element.

`target` 屬性指定了在哪裏打開鏈接文檔。 `target="_blank"` 意思是在新標籤頁或窗口中打開鏈接的文檔。

這是一個帶有 `target` 屬性的 `a` 元素的基本語法：

```html
<a href="https://www.freecodecamp.org" target="_blank">freeCodeCamp</a>
```

Add a `target` attribute with the value `_blank` to the anchor (`a`) element's opening tag, so that the link opens in a new tab.

# --hints--

Your `p` element should have a nested anchor (`a`) element with the text `cat photos`. You may have deleted it or have a typo.

```js
const anchor = document.querySelector('p > a');
assert.equal(
    anchor?.innerText?.toLowerCase().replace(/\s+/g, ' '), 'cat photos'
);
```

Your anchor (`a`) element does not have a `target` attribute. 請檢查在開始標籤的名稱後面要有一個空格，且/或所有的屬性名稱前面也要有一個空格。

```js
assert.isTrue(document.querySelector('a').hasAttribute('target'));
```

The value of the `target` attribute should be `_blank`. 你可能落下了這個值或者有拼寫錯誤。 所有屬性的值都應該放在引號中。

```js
assert.equal(document.querySelector('a').getAttribute('target'), '_blank');
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
--fcc-editable-region--
      <p>See more <a href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```


