---
id: 66984fe8efbd8c53a78c9b62
title: 步驟 37
challengeType: 0
dashedName: step-37
---

# --description--

The `title` element determines what browsers show in the title bar or tab for the page.

在 `head` 元素中添加一個具有以下文本的 `title` 元素。

`CatPhotoApp`

# --hints--

你刪除了 `head` 元素，或者它缺少開始標籤或結束標籤。

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

你的 `title` 元素應該嵌套在 `head` 元素中。 確保爲 `title` 元素添加開始標籤和結束標籤。

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

你的 `title` 元素應該有一個結束標籤。 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert.match(code, /\<\/title\>/);
```

你的 `title` 元素中的文本應該是 `CatPhotoApp`。 你可能忽略了文本或有拼寫錯誤。

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
--fcc-editable-region--
  <head>

  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <!-- TODO: Add link to cat photos -->
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>cat nip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Five cats looking around a field.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```


