---
id: 66984218720e29393b1adad2
title: 步驟 16
challengeType: 0
dashedName: step-16
---

# --description--

Before adding any new content, you should make use of a `section` element to separate the cat photos content from the future content.

`section` 元素用於在文檔中定義各部分，如章節、標題、頁腳或文檔的任何其他部分。 它是一個對 SEO 和無障礙特性有幫助的語義元素。

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Take your `h2`, comment, `p`, and anchor (`a`) elements and nest them in a `section` element.

# --hints--

Your `section` element should have an opening tag. 開始標籤的書寫語法爲：`<elementName>`。

```js
assert.exists(document.querySelector('section'));
```

Your `section` element should have a closing tag. 結束標籤在 `<` 字符後面要有一個 `/`。

```js
assert.match(code, /<\/section\s*>/i);
```

The entire `section` element should be between the opening and closing tags of the `main` element.

```js
assert.equal(document.querySelector('section').parentNode.nodeName, 'MAIN');
```

The existing `h2`, comment, `p` element, and anchor (`a`) element should be between the opening and closing tags of the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 3)
```

The `h1` element should not be nested in the `section` element.

```js
const childrenOfSection = [...document.querySelector('section').childNodes];
const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    --fcc-editable-region--
    <main>
      <h1>CatPhotoApp</h1>

      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>

    </main>
--fcc-editable-region--
  </body>
</html>
```


