---
id: 690ddb5e6ee94eb73ed172b0
title: Depure um Formulário de Doação
challengeType: 25
dashedName: debug-donation-form
# Purposefully removed the demo type because this is a debugging project.
---

# --description--

Uma instituição de caridade local criou um site de formulário de doação, mas há vários problemas que precisam ser corrigidos. O formulário não é acessível e possui alguns erros de sintaxe em HTML.

Seu trabalho é corrigir todos os erros para que o formulário funcione corretamente e seja acessível a todos os usuários. Complete os itens nas histórias de usuário abaixo e clique em "Run the Tests" para ver se você corrigiu todos os erros.

**Histórias de Usuário:**

1. Os elementos `input` são elementos void e não devem ter tags de fechamento. Remova todas as tags de fechamento `</input>` do formulário.
2. Adicione elementos `label` para cada campo de entrada do formulário para que os usuários saibam para que serve cada campo. O texto do label deve corresponder ao que está atualmente ao lado de cada entrada.
3. O tipo de entrada `Email Address:` deve ser um `email` em vez de `text`.
4. Você deve associar cada elemento `label` ao seu elemento `input` correspondente usando o atributo `for` no `label` e um atributo `id` correspondente no `input`.
5. Adicione o atributo `required` aos campos de entrada de texto, email e número (mas não à caixa de seleção ou ao botão de envio) para garantir que os usuários preencham as informações obrigatórias.

# --hints--

Você não deve ter nenhuma tag de fechamento `</input>` no seu código.

```js
assert.notMatch(code, /<\/input>/);
```

Você deve ter exatamente cinco elementos `input` no seu formulário.

```js
assert.lengthOf(document.querySelectorAll("input"), 5);
```

Você deve ter exatamente quatro elementos `label` no seu formulário.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Seu primeiro `label` deve ter o texto `Full Name:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[0]?.textContent, /Full\s*Name\s*:/i);
```

Seu primeiro `input` deve ter um atributo `required`.

```js
const firstInput = document.querySelector("input[name='name']");
assert.isTrue(firstInput?.hasAttribute("required"));
```

Seu segundo `label` deve ter o texto `Email Address:`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[1]?.textContent, /Email\s*Address\s*:/i);
```

Seu segundo `input` deve ter um atributo `required`.

```js
const secondInput = document.querySelector("input[name='email']");
assert.isTrue(secondInput?.hasAttribute("required"));
```

Seu terceiro `label` deve ter o texto `Donation Amount ($):`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[2]?.textContent, /Donation\s*Amount\s*\(\$\)\s*:/i);
```

Seu terceiro `input` deve ter um atributo `required`.

```js
const thirdInput = document.querySelector("input[name='amount']");
assert.isTrue(thirdInput?.hasAttribute("required"));
```

Seu quarto `label` deve ter o texto `Subscribe`.

```js
const labels = document.querySelectorAll("label");
assert.match(labels[3]?.textContent, /Subscribe/i);
```

Você deve associar cada elemento `input` a um elemento `label` usando os atributos `for` e `id`.

```js 
const els = document.querySelectorAll('input:not([type="submit"])'); 
assert.isNotEmpty(els); 
els.forEach(input => { 
  const label = document.querySelector(`label[for="${input.id}"]`); 
  assert.exists(label);  
  assert.equal(label.tagName, "LABEL"); 
})
``` 

Seu `input` para o endereço de email deve ser do tipo `email`.

```js
const emailInput = document.querySelector("input[name='email']");
assert.strictEqual(emailInput?.getAttribute("type"), "email");
```

Seu `input` de checkbox não deve ter um atributo `required`.

```js
const checkboxInput = document.querySelector("input[name='newsletter']");
assert.isFalse(checkboxInput?.hasAttribute("required"));
```

Seu `button` de envio não deve ter um atributo `required`.

```js
const submitButton = document.querySelector("input[type='submit']");
assert.isFalse(submitButton?.hasAttribute("required"));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    
    Full Name:
    <input type="text" name="name"></input>

    Email Address:
    <input type="text" name="email">

    Donation Amount ($):
    <input type="number" name="amount"></input>

    <input type="checkbox" name="newsletter"></input>
    Subscribe

    <input type="submit" value="Send"></input>
  </form>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <title>Donation Form</title>
</head>
<body>
  <h1>Donation Form</h1>
  <form>
    <label for="full-name">Full Name:</label>
    <input id="full-name" type="text" name="name" required>

    <label for="email">Email Address:</label>
    <input id="email" type="email" name="email" required>

    <label for="amount">Donation Amount ($):</label>
    <input id="amount" type="number" name="amount" required>

    <label for="newsletter">Subscribe</label>
    <input id="newsletter" type="checkbox" name="newsletter">

    <input type="submit" value="Send">
  </form>
</body>
</html>
```
