---
id: 66f53dc2c5bd6a11d6c3282f
title: بناء تطبيق مولد كلمات مرور
challengeType: 26
dashedName: lab-password-generator
---

# --description--

في هذا المختبر، ستتدرب على استخدام الدوال ببناء مولد كلمات مرور عشوائي.

**الهدف:** حقق قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تنشئ دالة تسمى `generatePassword` تأخذ معلمة تشير إلى طول كلمة المرور المولدة. يمكنك تسمية المعلمة كما تريد.
2. يجب أن تُرجع دالتك سلسلة نصية تمثل كلمة مرور مولدة عشوائيًا. استخدم السلسلة النصية التالية وطرق `Math` المختلفة لمساعدتك في إرجاع سلسلة جديدة تحتوي على حروف عشوائية: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. يجب أن تعرف متغيرًا يسمى `password` وتعيّن له نتيجة استدعاء دالة `generatePassword` مع معلمة رقمية تمثل طول كلمة المرور المطلوبة.
4. يجب أن يكون لديك `console.log` يسجل سلسلة واحدة تتكون من دمج الرسالة `Generated password:` مع المتغير `password` مفصولين بمسافة.

# --hints--

يجب أن يكون لديك دالة `generatePassword` بمعلمة. يمكنك تسمية المعلمة كما تريد.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

يجب أن تُرجع دالة `generatePassword` سلسلة نصية.

```js
const result = generatePassword(5);
assert.isString(result);
```

يجب أن تُرجع دالة `generatePassword` سلسلة جديدة بالطول الصحيح.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

يجب أن تُرجع دالتك كلمة مرور مولدة عشوائيًا تحتوي على حروف صحيحة.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

يجب أن تُرجع دالتك كلمة مرور مولدة عشوائيًا تحتوي على أكثر من حرف فريد واحد.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

يجب أن تُرجع دالتك سلسلة عشوائية جديدة في كل مرة تُستدعى فيها.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

يجب أن يكون لديك متغير `password`.

```js
assert.isDefined(password);
```

يجب أن يكون متغير `password` سلسلة نصية.

```js
assert.isString(password);
```

يجب أن تستدعي دالة `generatePassword` بمعلمة رقمية وتخزن كلمة المرور المرجعة في المتغير `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

يجب أن تسجل سلسلة واحدة تجمع بين `Generated password:` و`password` مفصولتين بمسافة واحدة باستخدام `+` أو قالب نصي.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
