---
id: 646c5ace05e4be4211407935
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Para ver el elemento `.cat-head`, dale un fondo de gradiente lineal con `#5e5e5e` al `85%` y `#45454f` al `100%`.

Es posible que no notes la diferencia entre estos dos colores, pero están ahí.

# --hints--

Tu selector `.cat-head` debe tener una propiedad `background`.

```js
assert.match(code, /background:/)
```

Tu propiedad `background` debe usar la función `linear-gradient`.

```js
assert.match(code, /background:\s*linear-gradient\(/)
```

Tu función `linear-gradient` debe establecer el primer color a `#5e5e5e` al `85%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const firstGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[1];
assert.equal(firstGradient,'(94,94,94)85%,');
```

Tu función `linear-gradient` debe establecer el segundo color a `#45454f` al `100%`.

```js
const gradientBackgroundImage = new __helpers.CSSHelp(document).getStyle('.cat-head')?.getPropVal('background-image', true);
const secondGradient = gradientBackgroundImage.split(/\s*rgb\s*/)[2];
assert.equal(secondGradient,'(69,69,79)100%)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--
.cat-head {
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;

}
--fcc-editable-region--
```
