---
id: 646ceb843412c74edee27a79
title: Paso 22
challengeType: 0
dashedName: step-22
---

# --description--

Ahora aprenderás un truco CSS para dibujar triángulos.

Se utilizará para dibujar las orejas del gato.

Mediante un selector de clase, asigna al elemento `.cat-right-ear` las propiedades `height` y `width` establecidas en `100px`. Establece el `background-color` a `white`. Establece los bordes izquierdo y derecho en `35px solid blue`. Establece los bordes superior e inferior en `35px solid red`.

# --hints--

Debes tener un selector `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Tu selector `.cat-right-ear` debe tener una propiedad `height` establecida en `100px`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height === '100px')
```

Tu selector `.cat-right-ear` debe tener una propiedad `width` establecida en `100px`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '100px')
```

Tu selector `.cat-right-ear` debe tener una propiedad `background-color` establecida en `white`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor === 'white')
```

Tu selector `.cat-right-ear` debe tener una propiedad `border-left` establecida en `35px solid blue`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid blue')
```

Tu selector `.cat-right-ear` debe tener la propiedad `border-right` establecida en `35px solid blue`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid blue')
```

Tu selector `.cat-right-ear` debería tener la propiedad `border-top` establecida en `35px solid red`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid red')
```

Tu selector `.cat-right-ear` debe tener la propiedad `border-bottom` establecida en `35px solid red`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--

--fcc-editable-region--
```
