---
id: 646c5d7057c45f432fcdd46c
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Usando un selector de clase, da al elemento `.cat-head` un ancho de `205px` y una altura de `180px`. Además, dale un borde de `1px solid #000` y un `border-radius` del `46%`.

# --hints--

Debes tener un selector `.cat-head`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head'))
```

Tu selector `.cat-head` debería tener un `width` establecido en `205px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.width === '205px');
```

Tu selector `.cat-head` debería tener un `height` con un valor de `180px`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.height === '180px')
```

Tu selector `.cat-head` debería tener un `border` con un valor de `1px solid #000`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.border === '1px solid rgb(0, 0, 0)')
```

Tu selector `.cat-head` debería tener un `border-radius` con un valor de `46%`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle(".cat-head")?.borderRadius === '46%')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  <main>
    <div class='cat-head'></div>
  </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

--fcc-editable-region--

--fcc-editable-region--
```
