---
id: 646c48df8674cf2b91020ecb
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Dentro de tu elemento `head`, añade una etiqueta `meta` con el atributo `charset` de `utf-8`. Añade también un elemento `title` con el texto `fCC Cat Painting`.

# --hints--

Deberías añadir un único elemento `meta`.

```js
assert(document.querySelectorAll('meta').length === 1);
```

Tu elemento `meta` debería tener un atributo `charset` con el valor establecido a `utf-8`.

```js
assert(document.querySelector('meta')?.getAttribute('charset')?.toLowerCase() === 'utf-8');
```

Deberías añadir un único elemento `title`.

```js
assert(document.querySelectorAll('title').length === 1);
```

Tu elemento `title` debería tener el texto `fCC Cat Painting`. Ten en cuenta tanto la ortografía como las mayúsculas y minúsculas.

```js
assert(document.querySelector('title')?.innerText === 'fCC Cat Painting');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    --fcc-editable-region--

    --fcc-editable-region--
</head>
<body></body>
</html>
```
