---
id: 646c59e942f35541923104bf
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Dentro de tu HTML, crea un elemento `main`. Dentro de ese elemento `main`, añade un elemento `div` con la clase `cat-head`.

# --hints--

Debe tener un elemento `main`.

```js
assert(document.querySelectorAll('main').length === 1);
```

Debe tener un elemento `div`.

```js
assert(document.querySelectorAll('div').length === 1);
```

Tu elemento `div` debe tener la clase `cat-head`.

```js
assert(document.querySelector('div')?.getAttribute('class') === 'cat-head');
```

Tu elemento `div` debe estar dentro de tu etiqueta `main`.

```js
assert(document.querySelectorAll('main div').length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
  --fcc-editable-region--

  --fcc-editable-region--
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}
```
