---
id: 6476fd4213318f6ee211028a
title: Paso 14
challengeType: 0
dashedName: step-14
---

# --description--

Ahora reemplaza el valor de la propiedad `position` de tu `.cat-head` con `fixed`. Deja `top` y `left` como están.

Después de eso, desplázate hacia arriba y abajo para ver cómo funciona el valor `fixed`.

# --hints--

Tu selector `.cat-head` debe tener la propiedad `position` establecida en `fixed`. Asegúrate de añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'fixed')
```

Tu selector `.cat-head` debe tener la propiedad `top` establecida en `100px`. Asegúrate de añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '100px')
```

Tu selector `.cat-head` debe tener la propiedad `left` establecida en `100px`. Asegúrate de añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '100px')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
      <div class="box"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

.box {
  width: 200px;
  height: 600px;
  background-color: #000;
  position: absolute;
  left: 650px;
  top: 800px;
}
```
