---
id: 6476f7a4827bcc61682f2347
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

La siguiente propiedad de posición es `absolute`. Cuando usas el valor `absolute` para tu propiedad `position`, el elemento se saca del flujo normal del documento, y luego su posición es determinada por las propiedades `top`, `right`, `bottom`, y `left`.

Establece la propiedad `position` de tu elemento `.cat-head` en `absolute` y, después, establece las propiedades `top` e `left` en cualquier valor de píxel positivo.

<!-- **Note**: You can experiment with `top`, `left`, `bottom`, and `right` properties here, but the test would only pass for `top` of `300px`, and left of `400px`. -->

# --hints--

Tu selector `.cat-head` debe tener una propiedad `position` establecida en `absolute`. Asegúrate de añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.position === 'absolute')
```

Tu selector `.cat-head` debe tener una propiedad `top` establecida en cualquier número en `px`. Asegúrate de añadir un punto y coma.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.top === '300px')
assert.match(code, /top:\s*\d{1,4}px;/)
```

Tu selector `.cat-head` debe tener una propiedad `left` establecida en cualquier número en `px`. Asegúrate de añadir un punto y coma.

```js
// assert(new __helpers.CSSHelp(document)?.getStyle('.cat-head')?.left === '400px')
assert.match(code, /left:\s*\d{1,4}px;/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head"></div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
--fcc-editable-region--
  position: relative;
  top: 100px;
  left: 100px;
--fcc-editable-region--
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}
```
