---
id: 6676a8b01e56ec1a1e07c202
title: Paso 23
challengeType: 0
dashedName: step-23
---

# --description--

Observa que ahora tienes un cuadrado blanco con bordes gruesos. Los bordes tienen esquinas diagonales que pueden utilizarse para crear triángulos.

Dentro del mismo selector de clase ajusta la `height` y la `width` a `0`. Establece los bordes izquierdo, derecho y superior como `transparent`.

Elimina la propiedad `background-color`, y debes ver un triángulo.

# --hints--


Debe tener un selector `.cat-right-ear`.

```js
assert(new __helpers.CSSHelp(document)?.getStyle('.cat-right-ear'))
```

Tu selector `.cat-right-ear` debe tener una propiedad `height` establecida en `0`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.height == '0px')
```

Tu selector `.cat-right-ear` debe tener una propiedad `width` establecida en `0`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.width === '0px')
```

Tu selector `.cat-right-ear` debe tener una propiedad `border-left` establecida en `35px solid transparent`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderLeft === '35px solid transparent')
```

Tu selector `.cat-right-ear` debe tener la propiedad `border-right` establecida en `35px solid transparent`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderRight === '35px solid transparent')
```

Tu selector `.cat-right-ear` debe tener una propiedad `border-top` establecida en `35px solid transparent`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderTop === '35px solid transparent')
```


Tu selector `.cat-right-ear` debe tener la propiedad `border-bottom` establecida en `35px solid red`. No olvides añadir un punto y coma.

```js
assert(new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.borderBottom === '35px solid red')
```

Tu selector `.cat-right-ear` no debe tener una propiedad `background-color`.

```js
assert(!new __helpers.CSSHelp(document).getStyle('.cat-right-ear')?.backgroundColor)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <title>fCC Cat Painting</title>
    <link rel="stylesheet" href="./styles.css">
</head>
<body>
    <main>
      <div class="cat-head">
        <div class="cat-ears">
          <div class="cat-left-ear">
            <div class="cat-left-inner-ear"></div>
          </div>
          <div class="cat-right-ear">
            <div class="cat-right-inner-ear"></div>
          </div>
        </div>
      </div>
    </main>
</body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  background-color: #c9d2fc;
}

.cat-head {
  position: absolute;
  right: 0;
  left: 0;
  top: 0;
  bottom: 0;
  margin: auto;
  background: linear-gradient(#5e5e5e 85%, #45454f 100%);
  width: 205px;
  height: 180px;
  border: 1px solid #000;
  border-radius: 46%;
}

--fcc-editable-region--
.cat-right-ear {
  height: 100px;
  width: 100px;
  background-color: white;
  border-left: 35px solid blue;
  border-right: 35px solid blue;
  border-top: 35px solid red;
  border-bottom: 35px solid red;
}
--fcc-editable-region--
```
