---
id: 699f38937b053e4679878ff1
title: 第 3 步
challengeType: 1
dashedName: step-3
---

# --description--

在使用原生 JavaScript 时，没有任何机制阻止你用字串或数组调用 `square` 函数，这会导致意想不到的结果。

为了说明这一点，将你的函数调用改为使用字串 `"something"`，而不是数字 `5`。

在下一课中，你将学习如何使用 TypeScript 为你的代码添加类型安全，防止这类误差。

# --hints--

你的 `square` 函数调用应该有一个 `"something"` 参数。

```js
const explorer = await __helpers.Explorer(code);
assert.isTrue(explorer.variables.result.value.matches("square('something')"));
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
