---
id: 699f4861dd5b5f164e1bda4a
title: 第 11 步
challengeType: 1
dashedName: step-11
---

# --description--

使用 TypeScript 的另一个好处是确保返回类型被强制执行且符合你的预期。有时 TypeScript 可以推导函数的返回类型，但你也可以显式指定它。

为说明这一点，将 `getAverage` 函数的返回类型更改为字串 `"something"`。

# --hints--

你的 `getAverage` 函数应该返回一个字串。

```js
assert.isString(getAverage([2, 14, 26, 8]));
```

你的 `getAverage` 函数应该返回字串 `"something"`。

```js
assert.strictEqual(getAverage([2, 14, 26, 8]), "something");
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  --fcc-editable-region--
  return sum / numbers.length;
  --fcc-editable-region--
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
