---
id: 699f412114c87221f16a99e3
title: 第 6 步
challengeType: 1
dashedName: step-6
---

# --description--

使用 TypeScript 的另一个好处是，如果你传递给函数调用的参数过多或过少，它会捕捉到误差。

为了说明这一点，请更新你的 `square` 函数调用，使其接受两个参数 `5` 和 `10`。

# --hints--

你的 `square` 函数调用应该有两个参数 `5` 和 `10`。

```js
const explorer = await __helpers.Explorer(code);
const { result } = explorer.variables;
assert.isTrue(result.value.matches("square(5, 10)"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5);
--fcc-editable-region--

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5, 10);

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
