---
id: 699f4639ae0016caaf81b7c6
title: 第 10 步
challengeType: 1
dashedName: step-10
---

# --description--

现在是测试你的函数的时候了。

首先创建一个名为 `avgResult` 的变量，并将 `getAverage` 的函数调用赋值给它。该函数调用的参数应该是 `[2, 14, 26, 8]`。

然后将 `avgResult` 的值输出到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个名为 `avgResult` 的变量。

```js
assert.isDefined(avgResult);
```

你应该将 `getAverage([2, 14, 26, 8])` 赋值给你的 `avgResult` 变量。

```js
assert.strictEqual(avgResult, getAverage([2, 14, 26, 8]));
```

你的 `avgResult` 变量应该是一个数字。

```js
assert.isNumber(avgResult);
```

你应该将 `avgResult` 输出到控制台。

```js
assert.equal(getLogs()[1], avgResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
