---
id: 699f513e01eb1768cc5df726
title: 第 17 步
challengeType: 1
dashedName: step-17
---

# --description--

既然 `exponent` 参数已标记为可选的，你现在可以只用一个参数调用该函数。

更新你的 `raiseTo` 函数调用，使其只传入数字 `3`。

# --hints--

你的 `raiseTo` 函数调用应只带一个参数 `3`。

```js
assert.match(
  code,
  /raiseTo\s*\(\s*3\s*\)/
);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}

--fcc-editable-region--
const raisedResult = raiseTo(2, 3);
--fcc-editable-region--
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
