---
id: 699f3738921fb79a651ded56
title: 第 2 步
challengeType: 1
dashedName: step-2
---

# --description--

现在是调用你的函数的时候了。

创建一个名为 `result` 的变量，并将 `square` 的函数调用赋值给它，传入参数 `5`。

在这下面，将 `result` 的值输出到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个名为 `result` 的变量。

```js
assert.isDefined(result);
```

你的 `result` 变量应该是一个数字。

```js
assert.isNumber(result);
```

你的 `result` 变量应该是 `5` 的平方。

```js
assert.strictEqual(result, 25);
```

你应该将 `result` 变量输出到控制台。

```js
assert.equal(getLogs()[0], result);
```

# --seed--

## --seed-contents--

```ts
function square(num) {
  return num * num;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
