---
id: 699f5063d2393c674b1f45dc
title: 第 16 步
challengeType: 1
dashedName: step-16
---

# --description--

正如你在之前的步骤中回顾的那样，如果你用错误数量的参数调用一个函数，TypeScript 会抛出误差。

有一种方法可以通过在参数名后添加 `?` 来将参数标记为可选的。这允许你用比参数更少的参数调用函数。

这是一个示例：

```ts
function greet(name: string, greeting?: string) {
  if (greeting) {
    return `${greeting}, ${name}!`;
  }
  return `Hello, ${name}!`;
}
```

在这个示例中，如果调用 `greeting` 函数时没有提供 `greet` 参数，它将是 `undefined`，函数会返回一个默认的问候语。

将你的 `raiseTo` 函数更新为使 `exponent` 参数成为可选的。

# --hints--

你的 `raiseTo` 函数应该将 `exponent` 参数标记为可选的。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[1].matches("exponent?: number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(2, 3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
