---
id: 699f1d72e906e4616155290d
title: 第 1 步
challengeType: 1
dashedName: step-1
---

# --description--

在本次研讨会中，你将继续学习类型注解以及它们如何与函数协作，通过创建一个类型安全的数学工具包来实现。

首先创建一个名为 `square` 的函数，该函数接受一个参数并返回一个数的平方。

# --hints--

你应该有一个名为 `square` 的函数。

```js
const explorer = await __helpers.Explorer(code);
assert.exists(explorer.allFunctions.square);
```

你的 `square` 函数应该有一个参数。

```js
const explorer = await __helpers.Explorer(code);
assert.lengthOf(explorer.allFunctions.square.parameters, 1);
```

你的 `square` 函数应该返回一个数的平方。

```js
assert.strictEqual(square(2), 4);
assert.strictEqual(square(3), 9);
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
