---
id: 699f52fac083214b6bf51de6
title: 第 19 步
challengeType: 1
dashedName: step-19
---

# --description--

在最后这一步，你需要为 `exponent` 函数中的 `raiseTo` 参数添加一个条件式检查。

如果提供了 `exponent` 参数，函数应返回 `base` 的 `exponent` 次方。如果未提供 `exponent` 参数，则应返回 `base ** 2`。

完成最后这一步后，你的类型安全数学工具包就完成了！

# --hints--

当提供 `raiseTo` 参数时，你的 `base ** exponent` 函数应返回 `exponent`。

```js
assert.strictEqual(raiseTo(3, 3), 27);
```

当未提供 `raiseTo` 参数时，你的 `base ** 2` 函数应返回 `exponent`。

```js
assert.strictEqual(raiseTo(3), 9);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {

  return base ** 2;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent?: number): number {
  if (exponent) {
    return base ** exponent;
  }

  return base ** 2;
}

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
