---
id: 699f3cb8d057f29954205d7d
title: 第 4 步
challengeType: 1
dashedName: step-4
---

# --description--

正如你从之前的课程中回忆的那样，你可以像这样为函数参数添加类型注解：

```ts
function add(num1: number, num2: number) {
  return num1 + num2;
}
```

在这个例子中，如果你尝试用非数字的参数调用 `add` 函数，TypeScript 会显示误差。

更新你的 `square` 函数，为 `num` 参数添加类型注解。类型应该是 `number`。

# --hints--

你应该为 `num` 函数中的 `square` 参数添加类型注解。类型应该是 `number`。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.square;
assert.isTrue(parameters[0].matches("num: number"));
```

# --seed--

## --seed-contents--

```ts
--fcc-editable-region--
function square(num) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square("something");

console.log(result);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
