---
id: 699f52220de88c8091032076
title: 第 18 步
challengeType: 1
dashedName: step-18
---

# --description--

如果你打开控制台，结果是 `NaN`，因为 `exponent` 的第二个参数从未提供。这个 `raiseTo` 函数需要做一些更新来修复这个问题。

第一个改动是返回 `base ** 2`。

在下一步中，你将看到如何处理可选的 `exponent` 参数。

# --hints--

你的 `raiseTo` 函数应该返回一个数字。

```js
assert.isNumber(raiseTo(4));
```

你的 `raiseTo` 函数应该返回 `base` 的 `2` 次幂。

```js
assert.strictEqual(raiseTo(4), 16);
assert.strictEqual(raiseTo(5), 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--
function raiseTo(base: number, exponent?: number): number {
  return base ** exponent;
}
--fcc-editable-region--

const raisedResult = raiseTo(3);
console.log(raisedResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
