---
id: 699f3f636aa840b7201aa37c
title: 第 5 步
challengeType: 1
dashedName: step-5
---

# --description--

如果你打开控制台，你应该会看到以下 TypeScript 错误信息：

```md
Argument of type 'string' is not assignable to parameter of type 'number'
```

现在参数被显式地指定为数字类型，TypeScript 能够捕捉到尝试向 `square` 函数传递字串参数的错误。

将传递给 `square` 函数的参数改为数字 `5`，以修复错误并在控制台中看到正确的输出。

# --hints--

你的 `square` 函数调用应该传递 `5` 作为参数，而不是 `"something"`。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square("something");
--fcc-editable-region--

console.log(result);
```
