---
id: 699f434e294a9e11cf4f65f0
title: 第 8 步
challengeType: 1
dashedName: step-8
---

# --description--

现在是为你的工具包项目添加另一个数学辅助函数的时候了。

首先创建一个名为 `getAverage` 的函数，它接受一个参数。你的函数应该返回一个数字数组的平均值。

在接下来的几个步骤中，你将为这个函数添加类型注解。

# --hints--

你应该有一个名为 `getAverage` 的函数。

```js
assert.isFunction(getAverage);
```

你的函数应该接受一个参数。

```js
assert.strictEqual(getAverage.length, 1);
```

你的函数应该返回一个数字。

```js
assert.isNumber(getAverage([1, 2, 3]));
```

你的函数应该对任何数字数组返回正确的结果。

```js
assert.strictEqual(getAverage([1, 2, 3, 4, 5]), 3);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
