---
id: 699f458e47fb88c756fdba81
title: 第 9 步
challengeType: 1
dashedName: step-9
---

# --description--

现在，`numbers` 参数具有隐式 `any` 类型。为了解决这个问题，你可以将该参数显式地指定为数字数组。

回想之前的课程，你可以这样为数字数组指定类型：

```ts
const numbers: number[] = [1, 2, 3];
```

更新你的 `getAverage` 函数，将 `numbers` 参数显式地指定为数字数组。

# --hints--

你的 `getAverage` 函数应该有一个名为 `numbers` 的参数，该参数显式地指定为数字数组。

```js
const explorer = await __helpers.Explorer(code);
const parameters = explorer.allFunctions.getAverage.parameters;
assert.isTrue(parameters[0].matches("numbers: number[]"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}
--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
