---
id: 699f4a3e7b4c6bfa538cc4f7
title: 第 13 步
challengeType: 1
dashedName: step-13
---

# --description--

如果你打开控制台，你会看到以下错误信息：

```md
Type 'string' is not assignable to type 'number'.
```

要解决这个误差，你需要将返回值改回 `sum / numbers.length`。

# --hints--

你的 `getAverage` 函数应该返回一个数字。

```js
assert.isNumber(getAverage([1, 2, 3]));
```

你的 `getAverage` 函数应该返回数组中数字的平均值。

```js
assert.strictEqual(getAverage([1, 2, 3]), 2);
assert.strictEqual(getAverage([10, 20, 30]), 20);
assert.strictEqual(getAverage([5, 10, 15, 20]), 12.5);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```
