---
id: 699f4b903bf5427a599f791d
title: 第 14 步
challengeType: 1
dashedName: step-14
---

# --description--

在接下来的几个步骤中，你将构建另一个函数，以继续练习函数的类型注解。

首先创建一个名为 `raiseTo` 的函数，它接受两个名为 `base` 和 `exponent` 的参数。两个参数都应该是 `number` 类型。该函数应返回 `base` 的 `exponent` 次幂的结果。函数的返回类型也应注解为 `number`。

# --hints--

你应该有一个 `raiseTo` 函数。

```js
assert.isFunction(raiseTo);
```

你的 `raiseTo` 函数应该有两个参数，`base` 和 `exponent`。

```js
assert.match(
  code,
   /function\s+raiseTo\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*|(?:const|let)\s+raiseTo\s*=\s*\(\s*base\s*(?::\s*number)?\s*,\s*exponent\s*(?::\s*number)?\s*\)\s*=>/
);
```

你的 `base` 和 `exponent` 参数都应该是 `number` 类型。

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.raiseTo;
assert.isTrue(parameters[0].matches("base: number"));
assert.isTrue(parameters[1].matches("exponent: number"));
```

你的 `raiseTo` 函数应该返回一个数字。

```js
assert.isNumber(raiseTo(2, 3));
```

你的 `raiseTo` 函数应该返回 `base` 的 `exponent` 次幂的结果。

```js
assert.strictEqual(raiseTo(2, 3), 8);
```

你的 `raiseTo` 函数应该有一个 `number` 的返回类型注解。

```js
const explorer = await __helpers.Explorer(code);
const { raiseTo } = explorer.allFunctions;
assert.isTrue(raiseTo.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
