---
id: 699f49669ea51a31d2ba89d8
title: 第 12 步
challengeType: 1
dashedName: step-12
---

# --description--

正如你从之前的课程中回忆的，你可以像这样显式地为函数指定返回类型：

```ts
function circleArea(radius: number): number {
  return Math.PI * radius * radius;
}
```

在这个例子中，`circleArea` 函数的返回类型是 `number`。如果你尝试返回不同的类型，TypeScript 会显示误差。

将你的 `getAverage` 函数更新为显式指定返回类型为 `number`。

# --hints--

你的 `getAverage` 函数应该有一个返回类型为 `number`。

```js
const explorer = await __helpers.Explorer(code);
const { getAverage } = explorer.allFunctions;
assert.isTrue(getAverage.hasReturnAnnotation("number"));
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

--fcc-editable-region--
function getAverage(numbers: number[]) {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}
--fcc-editable-region--

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```

# --solutions--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return "something";
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
