---
id: 699f41b8da20f816b5c7dece
title: 第 7 步
challengeType: 1
dashedName: step-7
---

# --description--

如果你打开控制台，你应该会看到以下错误信息：

```md
Expected 1 arguments, but got 2.
```

为了解决这个误差，更新你的 `square` 函数调用，使其只接受一个 `5` 类型的参数。

# --hints--

你的 `square` 函数调用应该只有一个 `5` 类型的参数。

```js
assert.strictEqual(result, 25);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

--fcc-editable-region--
const result = square(5, 10);
--fcc-editable-region--

console.log(result);
```
