---
id: 699f4ec939f10da62fb56c3e
title: 第 15 步
challengeType: 1
dashedName: step-15
---

# --description--

现在是测试你的函数的时候了。

创建一个名为 `raisedResult` 的变量，并将调用 `raiseTo` 并传入参数 `2` 和 `3` 的结果赋值给它。

然后，将 `raisedResult` 输出到控制台。

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
const getLogs = () => spy.calls.map(call => call?.[0]);
```

# --hints--

你应该有一个名为 `raisedResult` 的变量。

```js
assert.isDefined(raisedResult);
```

你的 `raisedResult` 变量应该是数字类型。

```js
assert.isNumber(raisedResult);
```

你应该将 `raisedResult` 输出到控制台。

```js
assert.equal(getLogs()[2], raisedResult);
```

# --seed--

## --seed-contents--

```ts
function square(num: number) {
  return num * num;
}

const result = square(5);

console.log(result);

function getAverage(numbers: number[]): number {
  const sum = numbers.reduce((acc, n) => acc + n, 0);
  return sum / numbers.length;
}

const avgResult = getAverage([2, 14, 26, 8]);

console.log(avgResult);

function raiseTo(base: number, exponent: number): number {
  return base ** exponent;
}

--fcc-editable-region--

--fcc-editable-region--
```

```json
{
  "compilerOptions": {
    "noCheck": true
  }
}
```
