---
id: 61537d86bdc3dd343688fceb
title: Крок 4
challengeType: 0
dashedName: step-4
---

# --description--

Далі додайте кожному `img` атрибут `src` відповідно до його порядку в документі. Перший `img` повинен мати `src` зі значенням `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg`. Решта повинні бути такими ж, за винятком заміни `1` на число, під яким `img` знаходиться в документі.

# --hints--

Всі дев’ять ваших `img` елементів повинні мати атрибут `src`.

```js
const images = [...document.querySelectorAll('img')];
assert.isTrue(images.every(image => image.getAttribute('src')));
```

Перший ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[0]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg');
```

Другий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[1]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg');
```

Третій ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[2]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg');
```

Четвертий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[3]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg');
```

П’ятий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[4]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg');
```

Шостий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[5]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg');
```

Сьомий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[6]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg');
```

Восьмий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[7]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg');
```

Дев’ятий ваш `img` елемент повинен мати `https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg` встановленим як значення атрибута `src`.

```js
assert.equal(document.querySelectorAll('img')?.[8]?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
--fcc-editable-region--
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
      <img>
--fcc-editable-region--
    </div>
  </body>
</html>
```

```css

```
