---
id: 61537c5f81f0cf325b4a854c
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Додайте `header` елемент всередину `body` елемента та призначте йому клас `header`.

Всередині `header` створіть `h1` з `css flexbox photo gallery` як текстом.

# --hints--

У вашому `header` елементі має бути `body` елемент.

```js
assert.exists(document.querySelector('body')?.querySelector('header'));
```

Ваш `header` елемент має мати `class` з `header` як значенням.

```js
assert.isTrue(document?.querySelector('body')?.querySelector('header')?.classList?.contains('header'));
```

Ваш `header` елемент має містити `h1` елемент.

```js
assert.exists(document.querySelector('.header')?.querySelector('h1'));
```

У вашому `h1` елементі має бути текст `css flexbox photo gallery`.

```js
assert.equal(document?.querySelector('.header')?.querySelector('h1')?.innerText, 'css flexbox photo gallery');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html>
  <head lang="en">
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
