---
id: 6153a3952facd25a83fe8083
title: Крок 21
challengeType: 0
dashedName: step-21
---

# --description--

Псевдоелемент `::after` створює елемент, який є останнім дочірнім елементом вибраного елемента. Ви можете використати його, щоб додати порожній елемент після останнього зображення. Якщо ви задасте йому той самий `width`, що й у зображень, це зсунеться останнє зображення вліво, коли галерея у двоколоночному макеті. Зараз воно по центру, бо ви встановили `justify-content: center` на flex-контейнер.

Приклад:

```CSS
.container::after {
  content: "";
  width: 860px;
}
```

Створіть новий селектор, використовуючи псевдоелемент `::after` на елементі `.gallery`. Додайте множину властивостей `content` зі значенням порожнього рядка `""` і `350px` для властивості `width`.

# --hints--

Ваш селектор `.gallery::after` має мати множину властивостей `""` зі значенням порожнього рядка для властивості `content`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.content, "\"\"");
```

Ваш селектор `.gallery::after` має мати `350px` для властивості `width`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery::after')?.width,  "350px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
* {
  box-sizing: border-box;
}

body {
  margin: 0;
  font-family: sans-serif;
  background: #f5f6f7;
}

.header {
  text-align: center;
  text-transform: uppercase;
  padding: 32px;
  background-color: #0a0a23;
  color: #fff;
  border-bottom: 4px solid #fdb347;
}

.gallery {
  display: flex;
  flex-direction: row;
  flex-wrap: wrap;
  justify-content: center;
  align-items: center;
  gap: 16px;
  max-width: 1400px;
  margin: 0 auto;
  padding: 20px 10px;
}

.gallery img {
  width: 100%;
  max-width: 350px;
  height: 300px;
  object-fit: cover;
  border-radius: 10px;
}

--fcc-editable-region--

--fcc-editable-region--
```
