---
id: 64dd076a1b33c86d84b4232f
title: Крок 5
challengeType: 0
dashedName: step-5
---

# --description--

Щоб краще уявити, як розмірюються ваші елементи, корисно додати кордон.

Додайте вашому `.gallery` елементу `width` зі значенням `50%` і множину `border`, встановлену на `5px solid red`.

Потім додайте вашим `img` елементам `width` зі значенням `100%`, множину `padding`, встановлену на `5px`, і множину `border`, встановлену на `5px solid blue`.

# --hints--

Ви повинні мати селектор `.gallery`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery'));
```

Ваш селектор `.gallery` повинен мати властивість `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.width);
```

Ваш селектор `.gallery` повинен мати властивість `width`, встановлену на `50%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.width, '50%');
```

Ваш селектор `.gallery` повинен мати властивість `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.gallery')?.border);
```

Ваш селектор `.gallery` повинен мати властивість `border`, встановлену на `5px solid red`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.gallery')?.border, '5px solid red');
```

Ви повинні мати селектор `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img'));
```

Ваш селектор `img` повинен мати властивість `width`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.width);
```

Ваш селектор `img` повинен мати властивість `width`, встановлену на `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.width, '100%');
```

Ваш селектор `img` повинен мати властивість `padding`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.padding);
```

Ваш селектор `img` повинен мати властивість `padding`, встановлену на `5px`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.padding, '5px');
```

Ваш селектор `img` повинен мати властивість `border`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('img')?.border);
```

Ваш селектор `img` повинен мати властивість `border`, встановлену на `5px solid blue`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('img')?.border, '5px solid blue');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Photo Gallery</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <header class="header">
      <h1>css flexbox photo gallery</h1>
    </header>
    <div class="gallery">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/1.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/2.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/3.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/4.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/5.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/6.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/7.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/8.jpg">
      <img src="https://cdn.freecodecamp.org/curriculum/css-photo-gallery/9.jpg">
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
