---
id: 68de7360e1832ea180577a8a
title: Einen Rechner für die n-te Fibonacci-Zahl erstellen
challengeType: 27
dashedName: build-an-nth-fibonacci-number-calculator
---

# --description--

**Ziel:** Erfüllen Sie die unten stehenden User Stories und bestehen Sie alle Tests, um das Lab abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `fibonacci` erstellen.
2. Sie sollten innerhalb der Funktion `sequence` eine Liste namens `fibonacci` definieren, die mit den Werten `[0, 1]` initialisiert wird.
3. Die Funktion `fibonacci` sollte einen Parameter akzeptieren, eine nicht-negative ganze Zahl `n`.  
4. Der Aufruf von `fibonacci(n)` sollte einen dynamischen Programmieransatz verwenden, um die `n`-te Zahl der Fibonacci-Folge zu berechnen und zurückzugeben, wobei jede Zahl die Summe der beiden vorhergehenden Zahlen ist.
5. Jede berechnete Fibonacci-Zahl sollte an die Liste `sequence` angehängt werden.

# --hints--

Sie sollten eine Funktion mit dem Namen `fibonacci` haben.

```js
({ test: () => runPython(`assert _Node(_code).has_function("fibonacci")`) })
```

Ihre Funktion `fibonacci` sollte einen einzelnen Parameter annehmen.

```js
({ test: () => runPython(`
from inspect import signature
sig = signature(fibonacci)
assert len(sig.parameters) == 1
`) })
```

Sie sollten innerhalb der Funktion `sequence` eine Liste namens `fibonacci` haben, die mit `[0, 1]` initialisiert wird.

```js
({ test: () => runPython(`assert _Node(_code).find_function("fibonacci").has_stmt("sequence = [0, 1]")`) })
```

`fibonacci(0)` sollte `0` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(0) == 0`) })
```

`fibonacci(1)` sollte `1` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(1) == 1`) })
```

`fibonacci(2)` sollte `1` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(2) == 1`) })
```

`fibonacci(3)` sollte `2` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(3) == 2`) })
```

`fibonacci(5)` sollte `5` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(5) == 5`) })
```

`fibonacci(10)` sollte `55` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(10) == 55`) })
```

`fibonacci(15)` sollte `610` zurückgeben.

```js
({ test: () => runPython(`assert fibonacci(15) == 610`) })
```

Sie sollten in Ihrem Code keine Rekursion verwenden.

```js
({ test: () => runPython(`
assert not _Node(_code).find_function("fibonacci").find_body().block_has_call("fibonacci")
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def fibonacci(n):
    sequence = [0, 1]
    
    if n < 2:
        return sequence[n]
    
    for i in range(2, n + 1):
        sequence.append(sequence[i - 1] + sequence[i - 2])
    
    return sequence[n]
```

