---
id: 656475bbf1c2573de1d2c69c
title: 步驟11
challengeType: 20
dashedName: step-11
---

# --description--

每次代碼運行時，你應該會看到 `all_characters` 字符串中的一個隨機字符。 這正是你要創建隨機密碼的目的。

不過，`random` 所依賴的算法使得生成的僞隨機數是可預測的。 因此，儘管 `random` 模塊適合於大多數常見的應用，因其決定性規律，它不能用於加密目的。

不要導入 `random`，而是導入 `secrets` 模塊。 然後將 `print()` 調用改爲使用 `secrets.choice(all_characters)`。

# --hints--

你應該導入 `secrets` 模塊而不是 `random` 模塊。

```js
({
  test: () => {
    assert.match(code, /^import\s+secrets/m)
    assert.isFalse( /^import\s+random/m.test(code))
  }
})
```

你應該修改你現有的 `print()` 調用，以打印 `secrets.choice(all_characters)`。

```js
({ test: () => assert.match(code, /^print\s*\(\s*secrets\.choice\s*\(\s*all_characters\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
print(random.choice(all_characters))
--fcc-editable-region--
```
