---
id: 656470d517833a39bb8b5608
title: 步驟 9
challengeType: 20
dashedName: step-9
---

# --description--

The `random` module contains a pseudo-random number generator. 它的大多數功能依賴於 `random()` 函數，該函數返回範圍在 `0.0`（包含）和 `1.0` （不包含）之間的浮點數。

調用 `random` 模塊中的 `random()` 函數並打印結果。

# --hints--

你應該打印調用 `random.random()` 的結果。

```js
({ test: () => assert.match(code, /^print\s*\(\s*random\.random\s*\(\s*\)\s*\)/m) })
```

# --seed--

## --seed-contents--

```py
import random
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation

# Combine all characters
all_characters = letters + digits + symbols

print(all_characters)
--fcc-editable-region--

--fcc-editable-region--
```
