---
id: 6564883669b5af4b69f794cc
title: 步驟6
challengeType: 20
dashedName: step-6
---

# --description--

你的 `all_characters` 變量是一個由所有小寫字母和大寫字母、所有 10 個數字和幾個特殊字符組成的字符串。

Just before it, add a comment saying `Combine all characters`.

# --hints--

You should add the comment just above your `all_characters` variable.

```js
({ test: () => assert.match(code, /^#\s*Combine all characters\s+^all_characters/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
import string


# Define the possible characters for the password
letters = string.ascii_letters
digits = string.digits
symbols = string.punctuation


all_characters = letters + digits + symbols
--fcc-editable-region--
```
