---
id: 67ef8f695db52583424bd118
title: 영화 리뷰 페이지 디자인하기
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `main` 요소가 있어야 합니다.
2. `main` 요소 안에 영화 제목을 위한 `h1` 요소가 있어야 합니다.
3. `h1` 요소 아래에 영화 표지를 표시하는 `img` 요소가 있어야 합니다. `img` 요소에는 이미지를 설명하는 `alt` 텍스트가 있어야 합니다. 원한다면 다음 이미지를 사용할 수 있습니다: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`.
4. 간단한 영화 설명을 담은 `p` 요소가 있어야 합니다.
5. 영화 평점을 표시하는 또 다른 `p` 요소가 있어야 합니다. 그 안에는 다음 항목들이 순서대로 있어야 합니다:
   - 텍스트 `strong`이 있는 `Movie Rating:` 요소.
   - `span` 속성이 `aria-hidden`로 설정된 `true` 요소로, 별표 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`를 사용해 평점을 시각적으로 나타냅니다.
   - 괄호 안에 평점을 나타내는 숫자 값(예: `9.2/10`)이 스팬 뒤에 있어야 합니다.
6. 텍스트 `h2`이 있는 `Cast Members` 요소가 있어야 합니다.
7. `ul` 요소가 있어야 합니다.
8. `ul` 요소 안에는 여러 `li` 요소가 있어야 하며, 각각은 배우 이름을 위한 `strong` 요소와 그 뒤에 텍스트 `as`가 앞에 붙은 해당 캐릭터 이름을 포함해야 합니다. (예: `James Holloway as Ethan Carter`).

# --hints--

`main` 요소가 있어야 합니다.

```js
assert.isNotNull(document.querySelector("main"));
```

`main` 요소 안에 영화 제목을 담은 `h1` 요소가 있어야 합니다.

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

`h1` 요소 아래에 설명용 `img` 속성이 있는 `alt` 요소가 있어야 합니다.

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

이미지 뒤에 영화 설명을 담은 `p` 요소가 있어야 합니다.

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

평점을 담은 또 다른 `p` 요소가 있어야 합니다.

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

두 번째 `p` 요소 안에 텍스트 `strong`가 있는 `Movie Rating: ` 요소가 있어야 합니다.

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

평점 단락 안에 `span` 요소가 있어야 합니다.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

평점 단락 안의 `span` 요소에는 채워진 별(⭐) 또는 빈 별(☆) 10개가 있어야 하며, `span` 뒤 괄호 안에 숫자 값이 있어야 합니다.

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

평점 단락 안의 `span` 요소에는 `aria-hidden` 속성이 `true`로 설정되어 있어야 합니다.

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

`h2` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector("h2"));
```

두 번째 `h2` 요소 뒤에 `p` 요소가 와야 합니다.

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

`h2` 요소에는 텍스트 `Cast Members`가 있어야 합니다.

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

`ul` 뒤에 `h2` 요소가 있어야 합니다.

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` 요소에는 여러 `li` 요소가 포함되어야 합니다.

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

각 `li` 요소는 배우 이름을 위한 `strong` 요소와 그 뒤에 텍스트 `as`가 앞에 붙은 해당 캐릭터 이름을 포함해야 합니다. (예: `James Holloway as Ethan Carter`).

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
