---
id: 5dc17dc8f86c76b9248c6eb5
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Los comentarios te permiten dejar mensajes sin afectar la visualización en el navegador. También te permiten desactivar código. Un comentario en HTML comienza con `<!--`, contiene cualquier número de líneas de texto y termina con `-->`.

Aquí tienes un ejemplo de un comentario con el texto `TODO: Remove h1`:

```html
<!-- TODO: Remove h1 -->
```

Añade un comentario sobre el elemento `p` con este texto:

`TODO: Add link to cat photos`

# --hints--

Tu comentario debe iniciar con `<!--`. Faltan uno o más de los caracteres que definen el inicio de un comentario.

```js
assert.match(code, /<!--/);
```

Tu comentario debe terminar con `-->`. Faltan uno o más de los caracteres que definen el final de un comentario.

```js
assert.match(code, /-->/);
```

Tu código no debe tener caracteres de apertura/cierre extra. Tienes los caracteres `<!--` o `-->` extra mostrándose en el navegador.

```js
const noSpaces = code.replace(/\s/g, '');
assert.isBelow(noSpaces.match(/<!--/g)?.length, 2)
assert.isBelow(noSpaces.match(/-->/g)?.length, 2);
```

Tu comentario debe estar encima del elemento `p`. Los tienes en el orden incorrecto.

```js
assert.match(
  code.replace(/\s/g, ''),
  /<!--(.*?)--><p>everyonelovescutecatsonline!<\/p>/i
);
```

Tu comentario debe contener el texto `TODO: Add link to cat photos`.

```js
assert.match(code, /<!--\s*todo:\s+add\s+link\s+to\s+cat\s+photos\s*-->/i);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <h1>CatPhotoApp</h1>
    <h2>Cat Photos</h2>
--fcc-editable-region--
    
--fcc-editable-region--
    <p>Everyone loves cute cats online!</p>
  </body>
</html>
```
