---
id: 5dc23f9bf86c76b9248c6eba
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Puedes agregar imágenes a tu sitio web usando el elemento `img`. Los elementos `img` tienen una etiqueta de apertura sin una etiqueta de cierre. Un elemento sin etiqueta de cierre se conoce como <dfn>elemento vacío</dfn>.

Agregue un elemento `img` debajo del elemento `p`. En este punto, ninguna imagen aparecerá en el navegador.

# --hints--

Su elemento `img` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('img'));
```

Tu elemento `img` no debe tener una etiqueta de cierre `</img>`.

```js
assert.notMatch(code, /<\/img\>/);
```

Debe tener solo un elemento `img`. Elimine cualquier extra.

```js
assert.lengthOf(document.querySelectorAll('img'), 1);
```

Su elemento `img` debe estar debajo del elemento `p`. Los tienes en el orden incorrecto.

```js
const collection = [...document.querySelectorAll('p,img')].map(
  (node) => node.nodeName
);
assert.isBelow(collection.indexOf('P'), collection.indexOf('IMG'));
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
