---
id: 5dc24073f86c76b9248c6ebb
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Los <dfn>atributos</dfn> de HTML son palabras especiales usadas dentro de la etiqueta de apertura de un elemento para controlar el comportamiento del elemento. El atributo `src` en un elemento `img` especifica la URL de la imagen (donde se encuentra la imagen).

Aquí hay un ejemplo de un elemento `img` con un atributo `src` apuntando al logo de freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Dentro del elemento `img` existente, agrega un atributo `src` con esta URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Tu código debe tener un elemento `img`. Puede que hayas eliminado el elemento `img` o no hayas rodeado el valor del atributo `src` con comillas.

```js
assert.exists(document.querySelector('img'));
```

Tu elemento `img` debe tener un atributo `src`. Has omitido el atributo o tienes un error tipográfico. Asegúrate de que haya un espacio entre el nombre del elemento y el nombre del atributo.

```js
assert.exists(document.querySelector('img')?.src);
```

Tu elemento `img` debe tener el atributo `src` establecido en '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Has omitido la URL o tienes un error tipográfico. El caso de la URL es importante.

```js
assert.equal(document.querySelector('img')?.src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Aunque has establecido el `src` de un elemento `img` a la URL correcta, se recomienda siempre rodear el valor de un atributo con comillas.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Tu elemento `img` debe cerrarse ya sea con `>` o `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```
