---
id: 5ddbd81294d8ddc1510a8e56
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

El texto de un enlace debe colocarse entre la etiqueta 
 de apertura y la etiqueta de cierre de un elemento anchor (`a`).

Aquí hay un ejemplo de un enlace con el texto `click here to go to freeCodeCamp.org`:

```html
<a href="https://www.freecodecamp.org">click here to go to freeCodeCamp.org</a>
```

Añade el texto `cat photos` al elemento anchor. Esto se convertirá en el texto del enlace.

# --hints--

Tu elemento (`a`) debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('a'));
```

Tu elemento anchor (`a`) debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/a\>/);
```

El texto de tu elemento anchor (`a`) debería ser `cat photos`. Asegúrate de poner el texto del enlace entre la etiqueta de apertura y etiqueta de cierre del elemento (`a`).

```js
assert.equal(
  document.querySelector('a')?.innerText?.trim().toLowerCase().replace(/\s+/g, ' '),
    'cat photos'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <a href="https://freecatphotoapp.com"></a>
--fcc-editable-region--
      <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back.">
    </main>
  </body>
</html>
```
