---
id: 5dfb655eeacea3f48c6300b3
title: Paso 25
challengeType: 0
dashedName: step-25
---

# --description--

El elemento `figure` representa contenido independiente y te permitirá asociar una imagen a una descripción.

Introduce la imagen que acabas de añadir dentro de un elemento `figure`.

# --hints--

Tu elemento `figure` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('figure'));
```

Tu elemento `figure` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/figure\>/);
```

Debe haber un elemento `figure` justo encima de la etiqueta de cierre del segundo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')[1]?.lastElementChild.nodeName, 'FIGURE');
```

El elemento `img` de lasaña debe estar anidado en el elemento `figure`.

```js
assert.exists(document.querySelector('figure > img'));
assert.equal(
    document.querySelector('figure > img').getAttribute('src').toLowerCase(),
      'https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg'
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
--fcc-editable-region--
        
        <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
        
--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```
