---
id: 5ef9b03c81a63668521804d4
title: Paso 34
challengeType: 0
dashedName: step-34
---

# --description--

El elemento `strong` se utiliza para indicar que una parte de un texto es importante o urgente.

En el `figcaption` que acabas de añadir, indica que `hate` tiene una fuerte importancia envolviéndolo en un elemento `strong`.

# --hints--

Tu elemento `strong` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('strong'));
```

Tu elemento `strong` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/strong\>/);
```

Tu elemento `strong` debe rodear la palabra `hate` en el texto `Cats hate other cats.`. O has omitido el texto o hay un error tipográfico.

```js
assert.equal(
  document
    .querySelectorAll('figcaption')[1]
    ?.querySelector('strong')
    ?.innerText?.trim().toLowerCase(), 'hate'
);
```

El texto del `figcaption` debe ser `Cats hate other cats.`. Verifica si hay errores tipográficos y que los espacios necesarios estén presentes alrededor de las etiquetas de apertura y cierre del elemento `strong`.

```js
const secondFigCaption = document.querySelectorAll('figcaption')[1];
assert.match(
   secondFigCaption?.innerText
      .replace(/\s+/gi, ' ')
      .trim(),
      /cats hate other cats\.?/i
);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
--fcc-editable-region--
        
          <figcaption>Cats hate other cats.</figcaption>  
        
--fcc-editable-region--
        </figure>
      </section>
    </main>
  </body>
</html>
```
