---
id: 5ef9b03c81a63668521804eb
title: Paso 39
challengeType: 0
dashedName: step-39
---

# --description--

El elemento `title` (título) determina lo que los navegadores muestran en la barra de título o en las pestañas del navegador.

Agrega un elemento `title` dentro del elemento `head` usando el texto a continuación: 

`CatPhotoApp`

# --hints--

Has eliminado el elemento `head` o le falta una etiqueta de apertura o de cierre.

```js
assert.match(code, /\<head\>/);
assert.match(code, /\<\/head\>/);
```

Tu elemento `title` debe estar anidado dentro del elemento `head`. Asegúrate de añadir una etiqueta de apertura y de cierre al elemento `title`.

```js
const noSpaces = code.replace(/\s/g, '');
assert.match(noSpaces, /\<head\>\<title\>.*\<\/title\>\<\/head\>/);
```

Tu elemento `title` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /\<\/title\>/);
```

El texto de tu elemento `title` deber ser `CatPhotoApp`. O has omitido el texto o tienes un error tipográfico.

```js
assert.equal(document.title?.toLowerCase(), 'catphotoapp');
```

# --seed--

## --seed-contents--

```html
<html>
  <head>
--fcc-editable-region--
    
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
        <h3>Top 3 things cats hate:</h3>
        <ol>
          <li>flea treatment</li>
          <li>thunder</li>
          <li>other cats</li>
        </ol>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/cats.jpg" alt="Two tabby kittens sleeping together on a couch.">
          <figcaption>Cats <strong>hate</strong> other cats.</figcaption>  
        </figure>
      </section>
    </main>
    <footer>
      <p>
        No Copyright - <a href="https://www.freecodecamp.org">freeCodeCamp.org</a>
      </p>
    </footer>
  </body>
</html>
```
