---
id: 5f07be6ef7412fbad0c5626b
title: Paso 18
challengeType: 0
dashedName: step-18
---

# --description--

Antes de añadir nuevo contenido, deberías utilizar un elemento `section`, para separar el contenido de cat photos, del contenido que añadiremos después.

El elemento `section` se utiliza para definir secciones en un documento, como capítulos, encabezados, pies de página o cualquier otra sección del documento. Es un elemento semántico que ayuda con el SEO y la accesibilidad.

```html
<section>
  <h2>Section Title</h2>
  <p>Section content...</p>
</section>
```

Toma tu elemento `h2`, dos elementos `p` y el elemento de ancla (`a`) y anídalos en un elemento `section`.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen la siguiente sintaxis: `<elementName>`.

```js
assert.exists(document.querySelector('section'));
```

Tu elemento `section` debe tener una etiqueta de cierre. Las etiquetas de cierre tiene una `/` después del carácter `<`.

```js
assert.match(code, /<\/section\s*>/i);
```

El elemento `section` debe estar entre las etiquetas de apertura y cierre del elemento `main`.

```js
assert.equal(document.querySelector('section')?.parentNode.nodeName, 'MAIN');
```

Los elementos existentes `h2`, `p` y el elemento de ancla (`a`) deben estar entre las etiquetas de apertura y cierre del elemento `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const foundElements = childrenOfSection.filter((child) => {
  return ['H2', 'A', 'P'].includes(child.nodeName);
});
assert.lengthOf(foundElements, 4)
```

El elemento `h1` no debe estar anidado en el elemento `section`.

```js
const sectionEl = document.querySelector('section');
assert.isNotNull(sectionEl);
const childrenOfSection = [...sectionEl.childNodes];

const includesH1 = childrenOfSection.some((child) => child.nodeName === 'H1');
assert.isFalse(includesH1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      
      <h2>Cat Photos</h2>
      <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
      <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
      <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      
--fcc-editable-region--
    </main>
  </body>
</html>
```
