---
id: 5f07c98cdb9413cbd4b16750
title: Paso 19
challengeType: 0
dashedName: step-19
---

# --description--

Es hora de añadir una nueva sección. Agrega un segundo elemento `section` debajo del elemento `section` existente.

# --hints--

Tu elemento `section` debe tener una etiqueta de apertura. Las etiquetas de apertura tienen esta sintaxis: `<elementName>`.

```js
assert.isAtLeast(document.querySelectorAll('section').length, 2);
```

Solo debes añadir una etiqueta de apertura `section`. Por favor, elimina cualquier adicional.

```js
assert.lengthOf(document.querySelectorAll('section'), 2);
```

Tu elemento `section` debe tener una etiqueta de cierre. Las etiquetas de cierre tienen un `/` justo después del carácter `<`.

```js
assert.isAtLeast(code.match(/<\/section>/g)?.length, 2);
```

Solo debes añadir una etiqueta de cierre `section`. Por favor, elimina cualquier adicional.

```js
assert.lengthOf(code.match(/<\/section>/g), 2);
```

El segundo elemento `section` no debe estar anidado en el primer elemento `section`.

```js
const el = document.querySelector('main > section');
assert.isNotNull(el);
const childrenOf1stSection = [
  ...el.children
];
const foundElems = childrenOf1stSection.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.isEmpty(foundElems);
```

Ambos elementos `section` deben estar entre las etiquetas de apertura y cierre del elemento `main`.

```js
const el = document.querySelector('main');
assert.isNotNull(el);
const childrenOfMain = [...el.children];
const foundElems = childrenOfMain.filter((child) => {
  return child.nodeName === 'SECTION';
});
assert.lengthOf(foundElems, 2);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
--fcc-editable-region--
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>

--fcc-editable-region--
    </main>
  </body>
</html>
```
