---
id: 68e76b1ce8216249d04df670
title: Construye una Clase Planet
challengeType: 27
dashedName: build-a-planet-class
---

# --description--

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes crear una clase llamada `Planet`.
1. La clase `Planet` debe tener un método `__init__` que:
   - Tiene cuatro parámetros: `self`, `name`, `planet_type` y `star`.
   - Lanza un `TypeError` con el mensaje `name, planet type, and star must be strings` si alguno de los argumentos pasados no es de tipo string.
   - Lanza un `ValueError` con el mensaje `name, planet_type, and star must be non-empty strings` si alguno de los argumentos pasados es una cadena vacía.
   - Asigna los valores pasados a los atributos de la instancia `name`, `planet_type` y `star`.
1. La clase `Planet` debería tener un método `orbit` que devuelva una cadena en el formato `{name} is orbiting around {star}...`.
1. La clase `Planet` debe tener un método `__str__` que devuelva una cadena en el formato `Planet: {name} | Type: {planet_type} | Star: {star}`.
1. Debes crear tres instancias de la clase `Planet` llamadas `planet_1`, `planet_2` y `planet_3`.
1. Debes imprimir cada objeto planeta para ver la salida del método `__str__`.
1. Debes llamar al método `orbit` en cada objeto planeta e imprimir el resultado.

# --hints--

Debes crear una clase llamada `Planet`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_class('Planet')
`
    )
})
```

La clase `Planet` debe tener un método `__init__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__init__')

`
    )
})
```

El método `__init__` debe tener cuatro parámetros: `self`, `name`, `planet_type` y `star`, en este orden.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').find_function('__init__').has_args('self, name, planet_type, star')

`
    )
})
```

El método `__init__` debe lanzar un `TypeError` con el mensaje `name, planet type, and star must be strings` si el primer argumento pasado no es de tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet(8, "planet", "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

El método `__init__` debe lanzar un `TypeError` con el mensaje `name, planet type, and star must be strings` si el segundo argumento pasado no es de tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", [], "sun")
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

El método `__init__` debe lanzar un `TypeError` con el mensaje `name, planet type, and star must be strings` si el tercer argumento pasado no es de tipo string.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", 0)
except TypeError as e:
  assert str(e) == "name, planet type, and star must be strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

El método `__init__` debe lanzar un `ValueError` con el mensaje `name, planet_type, and star must be non-empty strings` si el primer argumento pasado es una cadena vacía.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("", "planet", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

El método `__init__` debe lanzar un `ValueError` con el mensaje `name, planet_type, and star must be non-empty strings` si el segundo argumento pasado es una cadena vacía.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "", "sun")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

El método `__init__` debe lanzar un `ValueError` con el mensaje `name, planet_type, and star must be non-empty strings` si el tercer argumento pasado es una cadena vacía.

```js
({
    test: () => runPython(
`
try:
  planet = Planet("Name", "planet", "")
except ValueError as e:
  assert str(e) == "name, planet_type, and star must be non-empty strings"
else:
  assert False, "Expected to raise error..."
`
    )
})
```

Debes asignar `name` a `self.name` en el método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.name').is_equivalent('self.name = name')
`
    )
})
```

Debes asignar `planet_type` a `self.planet_type` en el método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.planet_type').is_equivalent('self.planet_type = planet_type')
`
    )
})
```

Debes asignar `star` a `self.star` en el método `__init__`.

```js
({
    test: () => runPython(
`
init = _Node(_code).find_class('Planet').find_function('__init__')
assert init.find_variable('self.star').is_equivalent('self.star = star')
`
    )
})
```

La clase `Planet` debe tener un método `orbit`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('orbit')
`
    )
})
```

El método `orbit` debería devolver una cadena en el formato `{name} is orbiting around {star}...`.

```js
({
    test: () => runPython(
`
pl = Planet('planetino', 'type', 'star')
assert pl.orbit() == 'planetino is orbiting around star...'
`
    )
})
```

La clase `Planet` debe tener un método `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).find_class('Planet').has_function('__str__')
`
    )
})
```

El método `__str__` debería devolver una cadena en el formato `Planet: {name} | Type: {planet_type} | Star: {star}`.

```js
({
    test: () => runPython(
`
assert str(Planet("abcd", "bcdea", "jhkl")) == "Planet: abcd | Type: bcdea | Star: jhkl"
`
    )
})
```

Debes crear tres instancias de la clase `Planet` llamadas `planet_1`, `planet_2` y `planet_3`.

```js
({
    test: () => runPython(
`
assert isinstance(planet_1, Planet)
assert isinstance(planet_2, Planet)
assert isinstance(planet_3, Planet)
`
    )
})
```

Debes imprimir cada objeto planeta para ver la salida del método `__str__`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1)')
assert _Node(_code).has_call('print(planet_2)')
assert _Node(_code).has_call('print(planet_3)')
`
    )
})
```

Debes llamar al método `orbit` en cada objeto planeta e imprimir el resultado.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_call('print(planet_1.orbit())')
assert _Node(_code).has_call('print(planet_2.orbit())')
assert _Node(_code).has_call('print(planet_3.orbit())')
`
    )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class Planet:
    def __init__(self, name, planet_type, star):
        args = (name, planet_type, star)
        if not all(isinstance(arg, str) for arg in args):
            raise TypeError("name, planet type, and star must be strings")
        if not all(args):
            raise ValueError('name, planet_type, and star must be non-empty strings')
        self.name = name
        self.planet_type = planet_type
        self.star = star   # The star the planet orbits

    def orbit(self):
        return f'{self.name} is orbiting around {self.star}...'

    def __str__(self):
        return f'Planet: {self.name} | Type: {self.planet_type} | Star: {self.star}'


# Test the Planet class
planet_1 = Planet('Earth', 'terrestrial', 'Sun')
planet_2 = Planet('Jupiter', 'gas giant', 'Sun')
planet_3 = Planet('Kepler-452b', 'super-Earth', 'Kepler-452')

print(planet_1.orbit()) # Earth is orbiting around Sun...
print(planet_1) # Planet: Earth | Type: terrestrial | Star: Sun

print(planet_2.orbit()) # Jupiter is orbiting around Sun...
print(planet_2) # Planet: Jupiter | Type: gas giant | Star: Sun

print(planet_3.orbit()) # Kepler-452b is orbiting around Kepler-452...
print(planet_3) # Planet: Kepler-452b | Type: super-Earth | Star: Kepler-452
```
