---
id: 68e51d8f163d3f8b70f9c4e5
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop, eserciterai l'uso dei metodi `includes()` e `slice()`.

Per cominciare, crea una variabile chiamata `fccSentence` e assegna a essa la stringa `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Dovresti dichiarare una variabile chiamata `fccSentence`.

```js
assert.exists(fccSentence);
```

La tua variabile `fccSentence` dovrebbe essere una stringa.

```js
assert.isString(fccSentence);
```

La tua variabile `fccSentence` dovrebbe essere assegnata alla stringa `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
