---
id: 68e6172975008fea742a63cf
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

Ora è il momento di registrare il risultato dell’uso del metodo `slice()`.

Usa un letterale template o la concatenazione di stringhe per registrare il messaggio `The word "<platform>" was sliced from the message.` nella console. Sostituisci `<platform>` con il valore effettivo della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti registrare il messaggio `The word "<platform>" was sliced from the message.`, dove `<platform>` deve essere sostituito con il valore effettivo della variabile.

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
