---
id: 68e6172975008fea742a63cd
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

Ora esplorerai il metodo `slice()`, che estrae una porzione di una stringa.

Crea una variabile chiamata `message` e assegnale la stringa `"Welcome to freeCodeCamp!"`.

Poi usa `console.log()` per stampare `"Here are some examples of the slice() method:"` nella console.

# --hints--

Dovresti dichiarare una variabile chiamata `message`.

```js
assert.exists(message);
```

La tua variabile `message` dovrebbe essere una stringa.

```js
assert.isString(message);
```

La tua variabile `message` dovrebbe essere `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Dovresti stampare `"Here are some examples of the slice() method:"` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
