---
id: 68e6172975008fea742a63d0
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

Puoi usare `slice()` per estrarre dall'inizio di una stringa.

Crea una variabile chiamata `greetingWord`. Poi assegnale il risultato dell'uso di `message.slice(0, 7)` per estrarre la prima parola `"Welcome"`.

Poi usa un template literal o la concatenazione di stringhe per stampare il messaggio `The first word is "<greetingWord>".` nella console. Sostituisci `<greetingWord>` con il valore effettivo della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti dichiarare una variabile chiamata `greetingWord`.

```js
assert.exists(greetingWord);
```

La tua variabile `greetingWord` dovrebbe essere una stringa.

```js
assert.isString(greetingWord);
```

Dovresti usare `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Dovresti assegnare il risultato di `message.slice(0, 7)` alla tua variabile `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

La tua variabile `greetingWord` dovrebbe essere uguale a `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Dovresti stampare il messaggio `The first word is "<greetingWord>".`, dove `<greetingWord>` dovrebbe essere sostituito dal valore effettivo della variabile.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
