---
id: 68e6172975008fea742a63cc
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

Il metodo `includes()` fa distinzione tra maiuscole e minuscole, quindi distingue tra lettere maiuscole e minuscole.

Crea una variabile chiamata `hasLowercaseFCC`. Poi assegnale il risultato dell'uso del metodo `includes()` per verificare se `fccSentence` include `"freecodecamp"` (tutto minuscolo).

Poi usa un template literal o la concatenazione di stringhe per stampare il messaggio `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` nella console. Sostituisci `<hasLowercaseFCC>` con il valore reale della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti dichiarare una variabile chiamata `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Dovresti usare `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Dovresti assegnare il risultato di `fccSentence.includes("freecodecamp")` alla tua variabile `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Il risultato di `fccSentence.includes("freecodecamp")` dovrebbe essere `false`.

```js
assert.isFalse(hasLowercaseFCC);
```

Dovresti stampare il messaggio che include `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, dove `<hasLowercaseFCC>` deve essere sostituito con il valore reale della variabile.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
