---
id: 68e6172975008fea742a63cb
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora è il momento di vedere cosa succede quando la sottostringa NON viene trovata.

Crea una variabile chiamata `hasJavaScript`. Poi assegnale il risultato dell'uso del metodo `includes()` per verificare se `fccSentence` include `"JavaScript"`.

Poi usa un template literal o la concatenazione di stringhe per stampare il messaggio `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` nella console. Sostituisci `<hasJavaScript>` con il valore effettivo della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti dichiarare una variabile chiamata `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Dovresti usare `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Dovresti assegnare il risultato di `fccSentence.includes("JavaScript")` alla tua variabile `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Il risultato di `fccSentence.includes("JavaScript")` dovrebbe essere `false`.

```js
assert.isFalse(hasJavaScript);
```

Dovresti stampare il messaggio `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, dove `<hasJavaScript>` deve essere sostituito con il valore effettivo della variabile.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
