---
id: 68e6172975008fea742a63d1
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

Il metodo `slice()` può usare indici negativi per contare dalla fine della stringa. `-1` indica l'ultimo carattere.

Ecco un esempio:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Crea una variabile chiamata `endPunctuation`. Poi assegnale il risultato dell'uso di `message.slice(-1)` per estrarre l'ultimo carattere.

Poi usa un template literal o la concatenazione di stringhe per stampare il messaggio `The ending punctuation mark is a "<endPunctuation>"` nella console. Sostituisci `<endPunctuation>` con il valore effettivo della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti dichiarare una variabile chiamata `endPunctuation`.

```js
assert.exists(endPunctuation);
```

La tua variabile `endPunctuation` dovrebbe essere una stringa.

```js
assert.isString(endPunctuation);
```

Dovresti usare `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Dovresti assegnare il risultato di `message.slice(-1)` alla tua variabile `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

La tua variabile `endPunctuation` dovrebbe essere uguale a `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Dovresti stampare il messaggio `The ending punctuation mark is a "<endPunctuation>"`, dove `<endPunctuation>` deve essere sostituito con il valore effettivo della variabile.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
