---
id: 68e6172875008fea742a63c9
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Come ricordi dalle lezioni precedenti, il metodo `includes()` verifica se una stringa contiene una sottostringa specifica e restituisce `true` o `false`.

Ecco un esempio che usa il metodo `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Crea una variabile chiamata `hasFreeCodeCamp`. Poi assegnale il risultato dell'uso del metodo `includes()` su `fccSentence` per verificare se contiene `"freeCodeCamp"`.

# --hints--

Dovresti dichiarare una variabile chiamata `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Dovresti usare `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Dovresti assegnare il risultato di `fccSentence.includes("freeCodeCamp")` alla tua variabile `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Il risultato di `fccSentence.includes("freeCodeCamp")` dovrebbe essere `true`.

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
