---
id: 68e6172975008fea742a63ca
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora usa un template literal o la concatenazione di stringhe per stampare il messaggio `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` sulla console. Sostituisci `<hasFreeCodeCamp>` con il valore reale della variabile.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Dovresti stampare il messaggio `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, dove `<hasFreeCodeCamp>` deve essere sostituito con il valore reale della variabile.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
