---
id: 68e6172975008fea742a63ce
title: Passaggio 8
challengeType: 1
dashedName: step-8
---

# --description--

Il metodo `slice()` estrae una sezione di una stringa. Prende due argomenti: l'indice di inizio (incluso) e l'indice di fine (non incluso).

Ecco un esempio:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Gli indici delle stringhe partono da 0. Nella stringa `message`, il carattere all'indice 11 è `f` (l'inizio di "freeCodeCamp").

Crea una variabile chiamata `platform`. Poi assegnale il risultato dell'uso di `message.slice(11, 23)` per estrarre `"freeCodeCamp"`.

# --hints--

Dovresti dichiarare una variabile chiamata `platform`.

```js
assert.exists(platform);
```

La tua variabile `platform` dovrebbe essere una stringa.

```js
assert.isString(platform);
```

Dovresti usare `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Dovresti assegnare il risultato di `message.slice(11, 23)` alla tua variabile `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

La tua variabile `platform` dovrebbe essere uguale a `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
