---
id: 68599af28ead38c22590a89d
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

يجب أن توجد طريقة لربط الأزرار بالمحتوى الذي تمثله. يتم ذلك باستخدام الخاصية `aria-controls` على كل زر. لاحقًا، ستضيف المحتوى وخصائص `id` المقابلة لعناصر المحتوى.

ابدأ بإضافة الخاصية `aria-controls` إلى كل عنصر من عناصر `button`. يجب تعيين قيمة الخاصية `aria-controls` للزر الأول إلى `"panel-earth"`. يجب تعيين الخاصية `aria-controls` للزر الثاني إلى `"panel-saturn"`. يجب تعيين الخاصية `aria-controls` للزر الثالث إلى `"panel-mars"`.

# --hints--

يجب أن يحتوي كل عنصر من عناصر `button` على خاصية `aria-controls`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[aria-controls]");
assert.lengthOf(btns, 3);
```

يجب أن يحتوي عنصر `button` الأول على خاصية `aria-controls` بقيمة `"panel-earth"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-controls"), "panel-earth");
``` 

يجب أن يحتوي عنصر `button` الثاني على خاصية `aria-controls` بقيمة `"panel-saturn"`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.getAttribute("aria-controls"), "panel-saturn");
```

يجب أن يحتوي عنصر `button` الثالث على خاصية `aria-controls` بقيمة `"panel-mars"`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2];
assert.equal(thirdBtn.getAttribute("aria-controls"), "panel-mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab">🌍 Earth</button>
        <button role="tab">🪐 Saturn</button>
        <button role="tab">🔴 Mars</button>
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
