---
id: 68599633407374c0d574587b
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

الخطوة التالية هي إضافة الأزرار التي ستتولى مسؤولية التبديل بين الكواكب المختلفة وحقائقها.

داخل عنصر `tablist`، أضف ثلاثة عناصر `button`، كل منها يحتوي على خاصية `role` مضبوطة على `"tab"`. يجب أن يحتوي أول عنصر `button` على النص `🌍 Earth`. يجب أن يحتوي العنصر الثاني `button` على النص `🪐 Saturn`. يجب أن يحتوي العنصر الثالث `button` على النص `🔴 Mars`.

# --hints--

يجب أن يكون لديك ثلاثة عناصر `button` داخل عنصر `tablist`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button");
assert.lengthOf(btns, 3);
```

يجب أن تحتوي كل عناصر `button` على خاصية `role` مضبوطة على `"tab"`.

```js
const btns = document.querySelectorAll("div[role='tablist'] button[role='tab']");
assert.lengthOf(btns, 3);
```

يجب أن يحتوي أول عنصر `button` على النص `🌍 Earth`.

```js
const firstBtn = document.querySelector("[role='tab']");
assert.equal(firstBtn.textContent, "🌍 Earth");
```

يجب أن يحتوي ثاني عنصر `button` على النص `🪐 Saturn`.

```js
const secondBtn = document.querySelectorAll("[role='tab']")[1];
assert.equal(secondBtn.textContent, "🪐 Saturn")
```

يجب أن يحتوي ثالث عنصر `button` على النص `🔴 Mars`.

```js
const thirdBtn = document.querySelectorAll("[role='tab']")[2]
assert.equal(thirdBtn.textContent, "🔴 Mars");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        
        --fcc-editable-region--
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
