---
id: 68599d0fcd5f6ac313d05908
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

عندما يختار المستخدم تبويبًا، يجب تحديث الخاصية `aria-selected` إلى `"true"` لذلك التبويب و`"false"` لجميع التبويبات الأخرى. هذا يشير إلى التبويب النشط حاليًا ويضمن إمكانية الوصول لتقنيات المساعدة.

افتراضيًا، يجب اختيار التبويب الأول. يمكنك فعل ذلك بإضافة الخاصية `aria-selected` إلى أول عنصر `button` وتعيين قيمتها إلى `"true"`.

# --hints--

يجب أن يحتوي أول عنصر `button` على خاصية `aria-selected`.

```js
assert.isTrue(
  document.querySelectorAll("[role='tab']")[0].hasAttribute("aria-selected")
);
```

يجب أن تحتوي أول عنصر `button` على خاصية `aria-selected` بقيمة `"true"`.

```js
const firstBtn = document.querySelectorAll("[role='tab']")[0];
assert.equal(firstBtn.getAttribute("aria-selected"), "true");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <div class="tabs">
      <h2 id="tabs-title">Planets</h2>
      <div role="tablist" aria-labelledby="tabs-title">
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-earth">🌍 Earth</button>
        --fcc-editable-region--
        <button role="tab" aria-controls="panel-saturn">🪐 Saturn</button>
        <button role="tab" aria-controls="panel-mars">🔴 Mars</button>
      </div>
    </div>

    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
