---
id: 685db0bd1280154ae85bcf9f
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستبني قائمة تبويبات كواكب يمكن الوصول إليها باستخدام HTML وCSS وJavaScript. تم توفير قالب HTML الجاهز وCSS لك.

في هذه الخطوات الأولى، ستبني بقية HTML الخاصة بقائمة التبويبات.

ابدأ بإنشاء عنصر `div` مع تعيين صنف له إلى `"tabs"`. داخل هذا العنصر `div`، أضف عنصر `h2` مع خاصية `id` مضبوطة على `"tabs-title"` ومحتوى نصي `Planets`.

# --hints--

يجب أن يكون لديك عنصر `div`.

```js
assert.exists(document.querySelector("div"));
```

يجب أن يحتوي عنصر `div` على صنف مضبوط إلى `"tabs"`.

```js
assert.exists(document.querySelector("div.tabs"));
```

يجب أن يحتوي عنصر `div` على عنصر `h2` بداخله.

```js
assert.exists(document.querySelector(".tabs h2"));
```

يجب أن يحتوي عنصر `h2` على خاصية `id` مضبوطة إلى `"tabs-title"`.

```js
assert.equal(document.querySelector(".tabs h2").getAttribute("id"), "tabs-title");
```

يجب أن يحتوي عنصر `h2` على المحتوى النصي `Planets`.

```js
assert.equal(document.querySelector(".tabs h2").textContent, "Planets");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Planets Facts</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    --fcc-editable-region--
    
    --fcc-editable-region--
    
    <script src="./script.js"></script>
  </body>
</html>
```

```css
.tabs [role="tablist"] {
  display: flex;
  gap: 0.5rem;
  margin-bottom: 1rem;
}

[role="tab"] {
  padding: 0.5rem 1rem;
  background: #eee;
  border: 1px solid #ccc;
  cursor: pointer;
  font-weight: bold;
}

[role="tab"][aria-selected="true"] {
  background: #fff;
  border-bottom: 2px solid dodgerblue;
}

[role="tabpanel"] {
  border: 1px solid #ccc;
  padding: 1rem;
}
```
