---
id: 587d7fab367417b2b2512bd9
title: 為你的散點圖圓形添加標籤
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

你可以添加文字來創建散點圖中點的標籤。

目標是顯示 `dataset` 中每個項目的第一個（`x`）和第二個（`y`）欄位的逗號分隔值。

`text` 節點需要 `x` 和 `y` 屬性來定位在 SVG 上。在這個挑戰中，`y` 值（決定高度）可以使用 `circle` 的 `cy` 屬性所使用的相同值。`x` 值可以比 `circle` 的 `cx` 值稍大一些，這樣標籤才會可見。這會將標籤推到繪製點的右側。

# --instructions--

使用 `text` 元素標示散點圖上的每個點。標籤的文字應該是兩個值，以逗號和空格分隔。 例如，第一個點的標籤是 `34, 78`。將 `x` 屬性設為比 `circle` 上 `cx` 屬性使用的值多 5 單位。將 `y` 屬性以與 `circle` 上 `cy` 值相同的方式設置。

# --hints--

你的程式碼應該有 10 個 `text` 元素。

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

第一個標籤應該有 `34, 78` 的文字、`39` 的 `x` 值和 `422` 的 `y` 值。

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

第二個標籤應該有 `109, 280` 的文字、`114` 的 `x` 值，以及 `220` 的 `y` 值。

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

第三個標籤應該有 `310, 120` 的文字、`315` 的 `x` 值，以及 `380` 的 `y` 值。

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

第四個標籤應該有 `79, 411` 的文字、`84` 的 `x` 值，以及 `89` 的 `y` 值。

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

第五個標籤應該有 `420, 220` 的文字，`x` 值為 `425`，`y` 值為 `280`。

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

第六個標籤應該有 `233, 145` 的文字、`238` 的 `x` 值，以及 `355` 的 `y` 值。

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

第七個標籤應該有 `333, 96` 的文字、`338` 的 `x` 值，以及 `404` 的 `y` 值。

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

第八個標籤應該有 `222, 333` 的文字，`x` 值為 `227`，`y` 值為 `167`。

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

第九個標籤應該有 `78, 320` 的文字、`83` 的 `x` 值，以及 `180` 的 `y` 值。

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

第十個標籤應該有 `21, 123` 的文字、`26` 的 `x` 值，以及 `377` 的 `y` 值。

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
