---
id: 587d7fac367417b2b2512bdc
title: >-
  Use the d3.max and d3.min Functions to Find Minimum and Maximum Values in a
  Dataset
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

D3 的 `domain()` 和 `range()` 方法會根據資料為你的刻度設定該資訊。有幾種方法可以讓這更簡單。

通常當你設定網域時，你會想要使用資料集中最小和最大值。在大型資料集中手動尋找這些值可能會導致錯誤。

D3 有兩個方法 - `min()` 和 `max()` 用來傳回這些資訊。以下是一個範例：

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

資料集可能包含巢狀陣列，例如散佈圖範例中的 `[x, y]` 座標配對。在這種情況下，你需要告訴 D3 如何計算最大值和最小值。幸運的是，`min()` 和 `max()` 兩個方法都接受回呼函式。在此範例中，回呼函式的引數 `d` 是目前的內部陣列。回呼函式需要傳回內部陣列中你想要計算最大值或最小值的元素（`x` 或 `y` 值）。以下是如何使用陣列的陣列來尋找最小值和最大值的範例：

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` 的值會是 `1`。

# --instructions--

`positionData` 陣列包含 x、y 和 z 座標的子陣列。使用 D3 方法找出陣列中 z 座標（第三個值）的最大值，並將其存入 `output` 變數。

# --hints--

`h2` 中的文字應該是 `8`。

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

你的程式碼應該使用 `max()` 方法。

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
