---
id: 587d7fab367417b2b2512bd8
title: 為你的 Circle Elements 添加屬性
challengeType: 6
forumTopicId: 301471
dashedName: add-attributes-to-the-circle-elements
---

# --description--

上一個挑戰為 `dataset` 中的每個點創建了 `circle` 元素，並將它們附加到 SVG。但 D3 需要更多關於每個 `circle` 位置和大小的資訊，才能正確顯示它們。

SVG 中的 `circle` 有三個主要屬性。`cx` 和 `cy` 屬性是座標。它們告訴 D3 在 SVG 上將形狀的*中心*定位在哪裡。半徑（`r` 屬性）決定 `circle` 的大小。

就像 `rect` 的 `y` 座標一樣，`circle` 的 `cy` 屬性是從 SVG 的頂部測量，而不是從底部。

這三個屬性都可以使用回呼函式來動態設定它們的值。請記住，所有串鏈在 `data(dataset)` 之後的函式會對 `dataset` 中的每個項目執行一次。回呼函式中的 `d` 參數指的是 `dataset` 中目前的項目，該項目是每個點的陣列。你可以使用括號表示法，例如 `d[0]`，來存取該陣列中的值。

# --instructions--

為你的 `circle` 元素添加 `cx`、`cy` 和 `r` 屬性。`cx` 值應該是 `dataset` 中每個條款陣列的第一個數字。`cy` 值應該基於陣列中的第二個數字，但請確保圖表是正向顯示而非倒置。所有圓形的 `r` 值應該是 `5`。

# --hints--

你的程式碼應該有 10 個 `circle` 元素。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

第一個 `circle` 元素的 `cx` 值應為 `34`，`cy` 值應為 `422`，`r` 值應為 `5`。

```js
const circle1 = document.querySelector('circle');
assert.strictEqual(circle1?.getAttribute('cx'), '34');
assert.strictEqual(circle1?.getAttribute('cy'), '422');
assert.strictEqual(circle1?.getAttribute('r'), '5');
```

第二個 `circle` 元素的 `cx` 值應為 `109`，`cy` 值應為 `220`，`r` 值應為 `5`。

```js
const circle2 = document.querySelectorAll('circle')[1];
assert.strictEqual(circle2?.getAttribute('cx'), '109');
assert.strictEqual(circle2?.getAttribute('cy'), '220');
assert.strictEqual(circle2?.getAttribute('r'), '5');
```

第三個 `circle` 元素的 `cx` 值應為 `310`，`cy` 值應為 `380`，`r` 值應為 `5`。

```js
const circle3 = document.querySelectorAll('circle')[2];
assert.strictEqual(circle3?.getAttribute('cx'), '310');
assert.strictEqual(circle3?.getAttribute('cy'), '380');
assert.strictEqual(circle3?.getAttribute('r'), '5');
```

第四個 `circle` 元素的 `cx` 值應為 `79`，`cy` 值應為 `89`，`r` 值應為 `5`。

```js
const circle4 = document.querySelectorAll('circle')[3];
assert.strictEqual(circle4?.getAttribute('cx'), '79');
assert.strictEqual(circle4?.getAttribute('cy'), '89');
assert.strictEqual(circle4?.getAttribute('r'), '5');
```

第五個 `circle` 元素的 `cx` 值應為 `420`，`cy` 值應為 `280`，`r` 值應為 `5`。

```js
const circle5 = document.querySelectorAll('circle')[4];
assert.strictEqual(circle5?.getAttribute('cx'), '420');
assert.strictEqual(circle5?.getAttribute('cy'), '280');
assert.strictEqual(circle5?.getAttribute('r'), '5');
```

第六個 `circle` 元素的 `cx` 值應為 `233`，`cy` 值應為 `355`，`r` 值應為 `5`。

```js
const circle6 = document.querySelectorAll('circle')[5];
assert.strictEqual(circle6?.getAttribute('cx'), '233');
assert.strictEqual(circle6?.getAttribute('cy'), '355');
assert.strictEqual(circle6?.getAttribute('r'), '5');
```

第七個 `circle` 元素的 `cx` 值應為 `333`，`cy` 值應為 `404`，而 `r` 值應為 `5`。

```js
const circle7 = document.querySelectorAll('circle')[6];
assert.strictEqual(circle7?.getAttribute('cx'), '333');
assert.strictEqual(circle7?.getAttribute('cy'), '404');
assert.strictEqual(circle7?.getAttribute('r'), '5');
```

第八個 `circle` 元素的 `cx` 值應為 `222`，`cy` 值應為 `167`，`r` 值應為 `5`。

```js
const circle8 = document.querySelectorAll('circle')[7];
assert.strictEqual(circle8?.getAttribute('cx'), '222');
assert.strictEqual(circle8?.getAttribute('cy'), '167');
assert.strictEqual(circle8?.getAttribute('r'), '5');
```

第九個 `circle` 元素的 `cx` 值應為 `78`，`cy` 值應為 `180`，`r` 值應為 `5`。

```js
const circle9 = document.querySelectorAll('circle')[8];
assert.strictEqual(circle9?.getAttribute('cx'), '78');
assert.strictEqual(circle9?.getAttribute('cy'), '180');
assert.strictEqual(circle9?.getAttribute('r'), '5');
```

第十個 `circle` 元素的 `cx` 值應為 `21`，`cy` 值應為 `377`，`r` 值應為 `5`。

```js
const circle10 = document.querySelectorAll('circle')[9];
assert.strictEqual(circle10?.getAttribute('cx'), '21');
assert.strictEqual(circle10?.getAttribute('cy'), '377');
assert.strictEqual(circle10?.getAttribute('r'), '5');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => d[0])
      .attr('cy', d => h - d[1])
      .attr('r', 5);
  </script>
</body>
```
