---
id: 587d7fa9367417b2b2512bd0
title: 反轉 SVG 元素
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

你可能注意到長條圖看起來像是顛倒或反轉的。這是因為 SVG 使用 (x, y) 座標的方式。

在 SVG 中，座標的原點位於左上角。`x` 座標為 0 時，形狀會放置在 SVG 區域的左邊緣。`y` 座標為 0 時，形狀會放置在 SVG 區域的上邊緣。較高的 `x` 值會將矩形推向右方。較高的 `y` 值會將矩形推向下方。

要讓長條圖正立，你需要改變 `y` 座標的計算方式。它需要同時考慮長條的高度和 SVG 區域的總高度。

SVG 區域的高度是 100。如果你的資料集中有一個資料點是 0，你會希望長條從 SVG 區域的底部開始（而不是頂部）。為此，`y` 座標需要一個 100 的值。如果資料點的值是 1，你會從 `y` 座標 100 開始，將長條設置在底部。然後你需要考慮長條高度為 1，因此最終的 `y` 座標會是 99。

`y = heightOfSVG - heightOfBar` 的 `y` 座標會將長條圖置於正立位置。

# --instructions--

將 `y` 屬性的回呼函式更改為設定長條圖正立。請記住，長條圖的 `height` 是資料值 `d` 的 3 倍。

**注意：** 一般來說，關係式為 `y = h - m * d`，其中 `m` 是縮放資料點的常數。

# --hints--

第一個 `rect` 的 `y` 值應該是 `64`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

第二個 `rect` 的 `y` 值應該是 `7`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

第三個 `rect` 的 `y` 值應該是 `34`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

第四個 `rect` 的 `y` 值應該是 `49`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

第五個 `rect` 的 `y` 值應該是 `25`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

第六個 `rect` 的 `y` 值應為 `46`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

第七個 `rect` 的 `y` 值應為 `13`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

第八個 `rect` 的 `y` 值應該是 `58`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

第九個 `rect` 的 `y` 值應該是 `73`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
