---
id: 587d7fa8367417b2b2512bcd
title: 為資料集中每個資料點創建一個 Bar
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

上一個挑戰只為 `svg` 元素添加了一個矩形來表現一個長條。在這裡，你將結合迄今為止所學的 `data()`、`enter()` 和 SVG 形狀，為 `dataset` 中的每個資料點創建並附加一個矩形。

先前的挑戰展示了如何為 `dataset` 中的每個條款創建並附加 `div` 的格式：

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

使用 `rect` 元素而非 `div` 元素時有一些差異。`rect` 元素必須附加到 `svg` 元素，而不是直接附加到 `body`。此外，你需要告訴 D3 在 `svg` 區域內放置每個 `rect` 的位置。長條圖的放置將在下一個挑戰中涵蓋。

# --instructions--

使用 `data()`、`enter()` 和 `append()` 方法為 `dataset` 中的每個項目創建並附加一個 `rect`。長條圖應該全部重疊顯示；這將在下一個挑戰中修正。

# --hints--

你的文件應該有 9 個 `rect` 元素。

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

你的程式碼應該使用 `data()` 方法。

```js
assert.match(code, /\.data/g);
```

你的程式碼應該使用 `enter()` 方法。

```js
assert.match(code, /\.enter/g);
```

你的程式碼應該使用 `append()` 方法。

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
