---
id: 587d7fa9367417b2b2512bce
title: 動態設定每個長條的座標
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

最後一個挑戰為 `dataset` 中的每個點創建並附加了一個矩形到 `svg` 元素，以表述一個長條。不幸的是，它們全部堆疊在彼此之上。

矩形的位置由 `x` 和 `y` 屬性控制。它們告訴 D3 在 `svg` 區域中從哪裡開始繪製形狀。上一個挑戰將它們各自設為 0，因此每個長條都放置在左上角。

對於長條圖，所有長條應該位於相同的垂直層級，這表示所有長條的 `y` 值保持相同（在 0）。然而，`x` 值需要隨著你添加新的長條而改變。請記住，較大的 `x` 值會將條款推得更靠右。當你遍歷 `dataset` 中的陣列項目時，`x` 值應該增加。

D3 中的 `attr()` 方法接受一個回呼函式來動態設定該屬性。回呼函式帶有兩個引數，一個是資料點本身（通常是 `d`），另一個是資料點在陣列中的索引。第二個索引引數是選用的。格式如下：

```js
selection.attr('property', (d, i) => {});
```

重要的是要注意，你不需要撰寫 `for` 迴圈或使用 `forEach()` 來迭代資料集中的項目。請記得，`data()` 方法會解析資料集，而任何在 `data()` 之後串鏈的行為都會對資料集中的每一個項目執行一次。

# --instructions--

更改 `x` 屬性回呼函式，使其傳回索引乘以 30。

**注意：** 每個長條寬度為 25，因此將每個 `x` 值增加 30 可在長條之間添加一些間距。在此範例中，任何大於 25 的值都適用。

# --hints--

第一個 `rect` 的 `x` 值應該是 `0`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

第二個 `rect` 的 `x` 值應該是 `30`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

第三個 `rect` 的 `x` 值應該是 `60`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

第四個 `rect` 的 `x` 值應該是 `90`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

第五個 `rect` 的 `x` 值應該是 `120`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

第六個 `rect` 的 `x` 值應該是 `150`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

第七個 `rect` 的 `x` 值應該是 `180`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

第八個 `rect` 的 `x` 值應該是 `210`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

第九個 `rect` 的 `x` 值應該是 `240`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
