---
id: 587d7fa8367417b2b2512bc9
title: 動態更新元素的高度
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

先前的挑戰涵蓋了如何從陣列中顯示資料以及如何添加 CSS 類別。你可以結合這些課程來創建一個簡單的長條圖。這有兩個步驟：

1. 為陣列中的每個資料點創建一個 `div`

2. 為每個 `div` 指定動態高度，使用 `style()` 方法中的回呼函式，將高度設為與資料值相等。

回想使用回呼函式設定樣式的格式：

```js
selection.style('cssProperty', d => d);
```

# --instructions--

為你的程式碼編輯器中的程式碼添加 `style()` 方法，以設定每個元素的 `height` 屬性。使用回呼函式傳回加上字串 `px` 的資料點值。

# --hints--

第一個 `div` 的 `height` 應該是 `12` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

第二個 `div` 應該有 `31` 像素的 `height`。

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

第三個 `div` 的 `height` 應該是 22 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

第四個 `div` 應該有 `height` 為 `17` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

第五個 `div` 應該有 `height` 為 `25` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

第六個 `div` 的 `height` 應為 `18` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

第七個 `div` 應該有 `29` 像素的 `height`。

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

第八個 `div` 的 `height` 應為 `14` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

第九個 `div` 的 `height` 應為 `9` 像素。

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
