---
id: 587d7fa6367417b2b2512bc3
title: 使用 D3 選取一個元素群組
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3 也有 `selectAll()` 方法來選取一組元素。它會傳回一個 HTML 節點的陣列，包含文件中所有符合輸入字串的項目。以下是一個選取文件中所有錨點標記的範例：

```js
const anchors = d3.selectAll('a');
```

像 `select()` 方法一樣，`selectAll()` 支援串鏈，並且你可以將它與其他方法一起使用。

# --instructions--

選取文件中所有的 `li` 頁籤，並透過串鏈 `.text()` 方法將它們的文字更改為字串 `list item`。

# --hints--

頁面上應該有 3 個 `li` 元素，且每個元素中的文字應該是 `list item`。大小寫和間距必須完全相符。

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

你的程式碼應該存取 `d3` 物件。

```js
assert.match(code, /d3/g);
```

你的程式碼應該使用 `selectAll` 方法。

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
