---
id: 587d7fac367417b2b2512bde
title: 使用預先定義的比例尺來放置元素
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

設定好刻度後，就該再次對散佈圖進行映射。刻度就像將 `x` 和 `y` 未經處理的資料轉換成適合且能正確在 SVG 上呈現的值的處理函式。它們會將資料保持在螢幕的繪圖區域內。

你使用縮放函式（scaling function）設定 SVG 形狀的座標屬性值。這包括 `rect` 或 `text` 元素的 `x` 和 `y` 屬性，或 `circles` 的 `cx` 和 `cy`。以下是一個範例：

```js
shape.attr('x', d => xScale(d[0]));
```

比例尺設定形狀座標屬性以將資料點放置到 SVG 上。當你顯示實際資料值時，例如在工具提示或標籤的 `text()` 方法中，則不需要套用比例尺。

# --instructions--

使用 `xScale` 和 `yScale` 將 `circle` 和 `text` 形狀定位到 SVG 上。對於 `circles`，應用比例尺來設定 `cx` 和 `cy` 屬性。也給它們一個半徑為 `5` 單位。

對於 `text` 元素，套用比例尺來設定 `x` 和 `y` 屬性。標籤應該偏移到點的右側。為此，將 `10` 單位添加到 `x` 資料值，然後再傳遞給 `xScale`。

# --hints--

你的程式碼應該有 10 個 `circle` 元素。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

第一個 `circle` 元素在應用縮放後，`cx` 值應約為 `91`，`cy` 值應約為 `368`。它的 `r` 值也應為 `5`。

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

第二個 `circle` 元素在應用縮放後，`cx` 值應約為 `159`，`cy` 值應約為 `181`。它的 `r` 值也應為 `5`。

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

第三個 `circle` 元素在應用縮放後，`cx` 值應約為 `340`，`cy` 值應約為 `329`。它的 `r` 值也應為 `5`。

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

第四個 `circle` 元素在套用比例後，`cx` 值應約為 `131`，`cy` 值應約為 `60`。它的 `r` 值也應為 `5`。

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

第五個 `circle` 元素在套用比例後，`cx` 值應約為 `440`，`cy` 值應約為 `237`。它的 `r` 值也應為 `5`。

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

第六個 `circle` 元素在套用比例後，`cx` 值應約為 `271`，`cy` 值應約為 `306`。它的 `r` 值也應為 `5`。

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

第七個 `circle` 元素在應用比例尺後，`cx` 值應約為 `361`，`cy` 值應約為 `351`。它的 `r` 值也應為 `5`。

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

第八個 `circle` 元素在應用縮放後，`cx` 值應約為 `261`，`cy` 值應約為 `132`。它的 `r` 值也應為 `5`。

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

第九個 `circle` 元素在套用比例後，`cx` 值應約為 `131`，`cy` 值應約為 `144`。它的 `r` 值也應為 `5`。

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

第十個 `circle` 元素在套用比例尺後，`cx` 值應約為 `79`，`cy` 值應約為 `326`。它的 `r` 值也應為 `5`。

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

你的程式碼應該有 10 個 `text` 元素。

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

第一個標籤在套用比例尺後，`x` 值應約為 `100`，`y` 值應約為 `368`。

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

第二個標籤在套用比例後，`x` 值應約為 `168`，`y` 值應約為 `181`。

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

第三個標籤在套用比例尺後，`x` 值應約為 `350`，`y` 值應約為 `329`。

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

第四個標籤在套用比例尺後，`x` 值應約為 `141`，`y` 值應約為 `60`。

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

第五個標籤在套用比例尺後，`x` 值應約為 `449`，`y` 值應約為 `237`。

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

第六個標籤在應用比例尺後，`x` 值應約為 `280`，`y` 值應約為 `306`。

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

第七個標籤在套用比例尺後，`x` 值應約為 `370`，`y` 值應約為 `351`。

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

第八個標籤在套用比例尺後，`x` 值應約為 `270`，`y` 值應約為 `132`。

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

第九個標籤在套用比例尺後，`x` 值應約為 `140`，`y` 值應約為 `144`。

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

第十個標籤在套用比例尺後，`x` 值應約為 `88`，`y` 值應約為 `326`。

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
