---
id: 587d7faa367417b2b2512bd6
title: 為你的 D3 元素添加工具提示
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

當使用者暫留在頁面上的某個條款時，工具提示會顯示該條款的更多資訊。有多種方式可以為視覺化添加工具提示。此挑戰使用 SVG 的 `title` 元素。

`title` 與 `text()` 方法配對，動態地將資料添加到長條圖。

# --instructions--

在每個 `rect` 節點下附加一個 `title` 元素。然後使用帶有回呼函式的 `text()` 方法，使文字顯示資料值。

# --hints--

你的程式碼應該有 9 個 `title` 元素。

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

第一個 `title` 元素應該有 `12` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

第二個 `title` 元素應該有 `31` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

第三個 `title` 元素應該有 `22` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

第四個 `title` 元素應該有 `17` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

第五個 `title` 元素應該有 `25` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

第六個 `title` 元素應該有 `18` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

第七個 `title` 元素應該有 `29` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

第八個 `title` 元素應該有 `14` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

第九個 `title` 元素應該有 `9` 的工具提示文字。

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
