---
id: 587d7fa8367417b2b2512bca
title: 變更長條圖的呈現方式
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

最後一個挑戰創建了一個長條圖，但有幾個格式上的更改可以改善它：

1. 為你的每個 bar 添加間距以視覺上分隔它們，這是透過為 `bar` 類別的 CSS 添加邊距來完成的。

2. 將數值乘以一個數字來縮放高度，以增加長條的高度，更清楚地顯示數值差異。

# --instructions--

首先，為 `style` 標記中的 `bar` 類別添加 `2px` 的 `margin`。接著，修改 `style()` 方法中的回呼函式，使其傳回原始資料值的 10 倍（加上 `px`）的值。

**注意：** 將每個資料點乘以*相同*的常數只會改變比例。這就像放大，並不會改變底層資料的意義。

# --hints--

第一個 `div` 應該有 `height` 為 `120` 像素和 `margin` 為 `2` 像素。

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

第二個 `div` 應該有 `height` 為 `310` 像素和 `margin` 為 `2` 像素。

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

第三個 `div` 應該有 `height` 為 `220` 像素和 `margin` 為 `2` 像素。

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

第四個 `div` 應該有 `height` 為 `170` 像素和 `margin` 為 `2` 像素。

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

第五個 `div` 應該有 `height` 為 `250` 像素和 `margin` 為 `2` 像素。

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

第六個 `div` 應該有 `height` 為 `180` 像素和 `margin` 為 `2` 像素。

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

第七個 `div` 應該有 `height` 為 `290` 像素和 `margin` 為 `2` 像素。

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

第八個 `div` 應該有 `height` 為 `140` 像素和 `margin` 為 `2` 像素。

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

第九個 `div` 應該有 `height` 為 `90` 像素和 `margin` 為 `2` 像素。

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
