---
id: 587d7fa7367417b2b2512bc6
title: 為你的元素添加行內樣式
challengeType: 6
forumTopicId: 301475
dashedName: add-inline-styling-to-elements
---

# --description--

D3 讓你可以使用 `style()` 方法為動態元素添加行內 CSS 樣式。

`style()` 方法接受以逗號分隔的鍵值對組作為引數。以下是一個將選取的文字顏色設為藍色的範例：

```js
selection.style('color', 'blue');
```

# --instructions--

為你的程式碼編輯器中的程式碼添加 `style()` 方法，使所有顯示的文字具有 `font-family` 為 `verdana`。

# --hints--

你的 `h2` 元素應該有 `verdana` 的 `font-family`。

```js
const headingTwo = document.querySelector('h2');
assert.exists(headingTwo); 
const fontFamily = window.getComputedStyle(headingTwo)['font-family'];
assert.strictEqual(fontFamily, 'verdana');
```

你的程式碼應該使用 `style()` 方法。

```js
assert.match(code, /\.style/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('font-family', 'verdana');
  </script>
</body>
```
