---
id: 587d7fa7367417b2b2512bc8
title: 為你的 D3 添加類別
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

在 HTML 元素上使用大量行內樣式即使對較小的應用程式來說也很難管理。為元素添加類別並使用 CSS 規則一次性設定該類別的樣式會比較容易。D3 有 `attr()` 方法可以為元素添加任何 HTML 屬性，包括類別名稱。

`attr()` 方法的運作方式與 `style()` 相同。它接受以逗號分隔的值，並且可以使用回呼函式。以下是一個為選取項目添加 `container` 類別的範例：

```js
selection.attr('class', 'container');
```

請注意，當你需要添加類別時，`class` 參數將保持不變，只有 `container` 參數會改變。

# --instructions--

為你的程式碼添加 `attr()` 方法，並在 `div` 元素上放置 `bar` 類別。

# --hints--

你的 `div` 元素應該有一個 `bar` 的類別。

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

你的程式碼應該使用 `attr()` 方法。

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
