---
id: 587d7faa367417b2b2512bd2
title: 為你的 D3 元素添加標籤
challengeType: 6
forumTopicId: 301476
dashedName: add-labels-to-d3-elements
---

# --description--

D3 讓你使用 SVG 的 `text` 元素為圖表元素（例如長條）添加標籤。

像 `rect` 元素一樣，`text` 元素需要有 `x` 和 `y` 屬性，才能將它放置在 SVG 上。它也需要存取要顯示那些值的資料。

D3 讓你對如何標記你的長條圖擁有高度的控制元件層級。

# --instructions--

編輯器中的程式碼已經將資料綁定到每個新的 `text` 元素。首先，將 `text` 節點附加到 `svg`。接著，為 `x` 和 `y` 座標添加屬性。它們應該以與 `rect` 相同的方式計算，除了 `text` 的 `y` 值應該讓標籤比長條高出 3 個單位。最後，使用 D3 的 `text()` 方法將標籤設定為資料點的值。

**注意：** 若要讓標籤比長條圖高，請決定 `text` 的 `y` 值應該比長條圖的 `y` 值大 3 或小 3。

# --hints--

第一個 `text` 元素應該有一個 `12` 的標籤和一個 `y` 值為 `61`。

```js
assert.strictEqual(document.querySelectorAll('text')[0]?.textContent, '12');
assert.strictEqual(
  document.querySelectorAll('text')[0]?.getAttribute('y'),
  '61'
);
```

第二個 `text` 元素應該有一個 `31` 的標籤和一個 `y` 值為 `4`。

```js
assert.strictEqual(document.querySelectorAll('text')[1]?.textContent, '31');
assert.strictEqual(
  document.querySelectorAll('text')[1]?.getAttribute('y'),
  '4'
);
```

第三個 `text` 元素應該有一個 `22` 的標籤和一個 `31` 的 `y` 值。

```js
assert.strictEqual(document.querySelectorAll('text')[2]?.textContent, '22');
assert.strictEqual(
  document.querySelectorAll('text')[2]?.getAttribute('y'),
  '31'
);
```

第四個 `text` 元素應該有一個 `17` 的標籤和一個 `y` 值為 `46`。

```js
assert.strictEqual(document.querySelectorAll('text')[3]?.textContent, '17');
assert.strictEqual(
  document.querySelectorAll('text')[3]?.getAttribute('y'),
  '46'
);
```

第五個 `text` 元素應該有一個 `25` 的標籤和一個 `22` 的 `y` 值。

```js
assert.strictEqual(document.querySelectorAll('text')[4]?.textContent, '25');
assert.strictEqual(
  document.querySelectorAll('text')[4]?.getAttribute('y'),
  '22'
);
```

第六個 `text` 元素應該有一個 `18` 的標籤和一個 `y` 值為 `43`。

```js
assert.strictEqual(document.querySelectorAll('text')[5]?.textContent, '18');
assert.strictEqual(
  document.querySelectorAll('text')[5]?.getAttribute('y'),
  '43'
);
```

第七個 `text` 元素應該有一個 `29` 的標籤和一個 `y` 值為 10。

```js
assert.strictEqual(document.querySelectorAll('text')[6]?.textContent, '29');
assert.strictEqual(
  document.querySelectorAll('text')[6]?.getAttribute('y'),
  '10'
);
```

第八個 `text` 元素應該有一個 `14` 的標籤和一個 `y` 值為 55。

```js
assert.strictEqual(document.querySelectorAll('text')[7]?.textContent, '14');
assert.strictEqual(
  document.querySelectorAll('text')[7]?.getAttribute('y'),
  '55'
);
```

第九個 `text` 元素應該有一個 `9` 的標籤和一個 `y` 值為 70。

```js
assert.strictEqual(document.querySelectorAll('text')[8]?.textContent, '9');
assert.strictEqual(
  document.querySelectorAll('text')[8]?.getAttribute('y'),
  '70'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg.selectAll('text').data(dataset).enter();
    // Add your code below this line




    // Add your code above this line
  </script>
  <body></body>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3)
      .text(d => d);
  </script>
  <body></body>
</body>
```
