---
id: 587d7fa8367417b2b2512bcb
title: 了解 D3 中的 SVG
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> 代表 <dfn>Scalable Vector Graphics</dfn>。

這裡的「可縮放」是指，如果你對一個物件進行放大或縮小，它不會出現像素化。它會隨著顯示系統縮放，無論是在小型行動螢幕還是大型電視螢幕上。

SVG 用於製作常見的幾何圖形。由於 D3 將資料映射成視覺表現，它使用 SVG 來創建視覺化的圖形。網頁的 SVG 圖形必須放在 HTML 的 `svg` 頁籤內。

當樣式使用相對單位（例如 `vh`、`vw` 或百分比）時，CSS 可以具備可擴展性，但使用 SVG 在建置資料視覺化時更具彈性。

# --instructions--

為你的 `body` 使用 `append()` 添加一個 `svg` 節點。使用 `attr()` 或 `style()` 方法分別為它設定 `width` 屬性為提供的 `w` 常數，以及 `height` 屬性為提供的 `h` 常數。你會在輸出中看到它，因為在 `style` 頁籤中對它套用了粉紅色的 `background-color`。

**注意：** 使用 `attr()` 時，寬度和高度屬性不帶單位。這是縮放的建置區塊－無論縮放層級為何，元素的寬高比始終為 5:1。

# --hints--

你的文件應該有 1 個 `svg` 元素。

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

`svg` 元素應該有一個 `width` 屬性設定為 `500` 或樣式設定為寬度為 500px。

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

`svg` 元素應該有一個 `height` 屬性設定為 `100` 或以樣式設定高度為 100px。

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
