---
id: 587d7fac367417b2b2512bdb
title: 在 `Scale` 上設定一個 Domain 和一個 Range
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

預設情況下，比例尺使用恆等關係。這表示輸入值對應到輸出值。然而，比例尺可以更加靈活且有趣。

假設一個資料集的值範圍從 50 到 480。這是刻度的輸入資訊，也稱為 <dfn>domain</dfn>。

你想要將那些點映射到 SVG 的 `x` 軸上，介於 10 單位和 500 單位之間。這是輸出資訊，也稱為 <dfn>range</dfn>。

`domain()` 和 `range()` 方法為比例尺設定這些值。這兩個方法都接受一個至少包含兩個元素的陣列作為引數。以下是一個範例：

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

依序，以下數值會顯示在主控台：`10`、`500`、`323.37` 和 `807.67`。

請注意，刻度使用域和值域之間的線性關係來計算給定數字的輸出。域中的最小值（50）對應到範圍中的最小值（10）。

# --instructions--

創建一個 scale 並將其 domain 設定為 `[250, 500]`，range 設定為 `[10, 150]`。

**注意：**你可以將 `domain()` 和 `range()` 方法串鏈到 `scale` 變數上。

# --hints--

你的程式碼應該使用 `domain()` 方法。

```js
assert.match(code, /\.domain/g);
```

`scale` 的 `domain()` 應該設定為 `[250, 500]`。

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

你的程式碼應該使用 `range()` 方法。

```js
assert.match(code, /\.range/g);
```

`scale` 的 `range()` 應該設定為 `[10, 150]`。

```js
assert.deepEqual(scale.range(), [10, 150]);
```

`h2` 中的文字應該是 `-102`。

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
