---
id: 587d7fab367417b2b2512bd7
title: 使用 SVG 圓形創建散點圖
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

散佈圖是另一種視覺化型別。它通常使用圓圈來對應具有兩個值的資料點。這些值分別對應到 `x` 軸和 `y` 軸，並用於定位視覺化中的圓圈位置。

SVG 有一個 `circle` 頁籤用來創建圓形。它的運作方式與你用於長條圖的 `rect` 元素非常相似。

# --instructions--

使用 `data()`、`enter()` 和 `append()` 方法，將 `dataset` 綁定到新增並附加到 SVG 的 `circle` 元素。

**注意：** 這些圓形不會顯示，因為我們還沒有設定它們的屬性。 我們會在下一個挑戰中完成這個動作。

# --hints--

你的程式碼應該有 10 個 `circle` 元素。

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
