---
id: 587d7fa9367417b2b2512bcf
title: 動態變更每個長條的高度
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

每個長條的高度可以設定為陣列中資料點的值，類似於 `x` 值的動態設定方式。

```js
selection.attr('property', (d, i) => {});
```

這裡 `d` 將會是資料點的值，而 `i` 將會是該資料點在陣列中的索引。

# --instructions--

將 `height` 屬性的回呼函式更改為傳回資料值乘以 3。

**注意：** 請記得將所有資料點乘以相同的常數會縮放資料（就像放大一樣）。這有助於在此範例中看出長條值之間的差異。

# --hints--

第一個 `rect` 應該有 `36` 的 `height`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

第二個 `rect` 的 `height` 應該是 `93`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

第三個 `rect` 的 `height` 應該是 `66`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

第四個 `rect` 的 `height` 應該是 `51`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

第五個 `rect` 的 `height` 應該是 `75`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

第六個 `rect` 的 `height` 應為 `54`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

第七個 `rect` 的 `height` 應為 `87`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

第八個 `rect` 的 `height` 應為 `42`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

第九個 `rect` 的 `height` 應為 `27`。

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
