---
id: 587d7fad367417b2b2512bdf
title: 為你的 Visualization 添加 Axes
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

另一種改進散佈圖的方法是為你的 x 軸和 y 軸添加標籤。

D3 有兩個方法，`axisLeft()` 和 `axisBottom()`，分別用來呈現 y 軸和 x 軸。以下是一個範例，示範如何根據先前挑戰中的 `xScale` 來創建 x 軸：

```js
const xAxis = d3.axisBottom(xScale);
```

下一步是在 SVG 上呈現軸。為此，你可以使用一般的 SVG 組件，即 `g` 元素。`g` 代表群組。與 `rect`、`circle` 和 `text` 不同，軸在呈現時只是一條直線。因為它是簡單的形狀，所以使用 `g` 是可行的。最後一步是套用 `transform` 屬性，將軸定位在 SVG 上的正確位置。否則，該線條會沿著 SVG 的邊框呈現，且無法被看見。SVG 支援不同類型的 `transforms`，但定位軸需要使用 `translate`。當它套用到 `g` 元素時，會將整個群組依指定的數值向右和向下移動。以下是一個範例：

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

上述程式碼將 x 軸放置在 SVG 的底部。然後它作為引數傳遞給 `call()` 方法。y 軸的運作方式相同，除了 `translate` 引數的表單為 `(x, 0)`。因為 `translate` 是在上述 `attr()` 方法中的字串，你可以使用串接來包含其引數的變數值。

# --instructions--

散佈圖現在有一個 x 軸。使用 `axisLeft()` 方法，在名為 `yAxis` 的變數中創建一個 y 軸。然後使用 `g` 元素呈現該軸。請確保使用 `transform` 屬性將軸向右平移 padding 單位數，向下平移 0 單位。記得要 `call()` 該軸。

# --hints--

你的程式碼應該使用 `axisLeft()` 方法，並將 `yScale` 作為引數傳入。

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

y 軸的 `g` 元素應該有一個 `transform` 屬性，用來將軸平移 `(60, 0)`。

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

你的程式碼應該呼叫 `yAxis`。

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
