---
id: 587d7fa8367417b2b2512bcc
title: 使用 SVG 顯示形狀
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

上一個挑戰創建了一個具有指定寬度和高度的 `svg` 元素，因為在 `style` 頁籤中為它套用了 `background-color`，所以它是可見的。程式碼為指定的寬度和高度留出了空間。

下一步是創建一個形狀放入 `svg` 區域。SVG 中有多種支援的形狀，例如矩形和圓形。它們用來顯示資料。例如，一個矩形（`<rect>`）SVG 形狀可以創建長條圖中的一個條。

當你將形狀放入 `svg` 區域時，你可以使用 `x` 和 `y` 座標來指定它的位置。原點 (0, 0) 位於左上角。`x` 的正值會將形狀向右推移，`y` 的正值會將形狀從原點向下推移。

要將形狀置於上一個挑戰中 500（寬）x 100（高）的 `svg` 中間，`x` 座標應為 250，`y` 座標應為 50。

SVG 的 `rect` 有四個屬性。它有 `x` 和 `y` 座標來指定它在 `svg` 區域中的位置。它也有 `height` 和 `width` 來指定大小。

# --instructions--

使用 `append()` 為 `svg` 添加一個 `rect` 形狀，並給它一個 `width` 屬性為 `25` 和 `height` 屬性為 `100`。同時，給 `rect` 的 `x` 和 `y` 屬性各設為 `0`。

# --hints--

你的文件應該有 1 個 `rect` 元素。

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

`rect` 元素應該有一個 `width` 屬性設定為 `25`。

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

`rect` 元素應該有一個 `height` 屬性設定為 `100`。

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

`rect` 元素應該有一個 `x` 屬性設為 `0`。

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

`rect` 元素應該有一個 `y` 屬性設定為 `0`。

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
