---
id: 587d7fac367417b2b2512bdd
title: 使用動態刻度
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

D3 的 `min()` 和 `max()` 方法對於設定刻度非常有用。

給定一組複雜的資料，首要任務是設定比例，使視覺化符合 `SVG` 容器的寬度和高度。你希望所有資料都繪製在 `SVG` 內，以便在網頁上可見。

以下範例設定散佈圖資料的 x 軸刻度。`domain()` 方法將關於繪圖未經處理的資料值的資訊傳遞給刻度。`range()` 方法則提供關於網頁上視覺化實際空間的資訊。

在範例中，定義域從 0 到集合中的最大值。它使用帶有基於陣列中 x 值的回呼函式的 `max()` 方法。值域使用 SVG 的寬度（`w`），但也包含一些內距。這會在散點圖的點與 SVG 邊緣之間留出空間。

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

填充一開始可能會讓人感到困惑。想像 x 軸是一條從 0 到 500（SVG 的寬度值）的水平線。將填充包含在 `range()` 方法中會強制繪圖從該線上的 30 開始（而不是 0），並在 470 結束（而不是 500）。

# --instructions--

使用 `yScale` 變數來創建線性 y 軸刻度。定義域應該從零開始，直到集合中最大的 `y` 值。區間應該使用 SVG 高度（`h`）並包含內邊距。

**注意：** 請記得保持圖表正立。當你設定 y 座標的範圍時，較高的值（高度減去內距）是第一個引數，較低的值是第二個引數。

# --hints--

`h2` 中的文字應為 `30`。

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

yScale 的 `domain()` 應該等價於 `[0, 411]`。

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

`yScale` 的 `range()` 應該等價於 `[470, 30]`。

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
