---
id: 587d7fa7367417b2b2512bc5
title: 在 D3 中處理動態資料
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

最後兩個挑戰涵蓋了使用 `data()` 和 `enter()` 方法以動態方式顯示資料的基礎。這些方法會取得一組資料，並且與 `append()` 方法一起，為資料集中的每個項目建立新的 DOM 元素。

在前一個挑戰中，你為 `dataset` 陣列中的每個條款創建了一個新的 `h2` 元素，但它們都包含相同的文字，`New Title`。這是因為你尚未使用綁定到每個 `h2` 元素的資料。

D3 的 `text()` 方法可以接受字串或回呼函式作為引數：

```js
selection.text(d => d);
```

在上述範例中，參數 `d` 取用綁定選取的資料集中的單一項目。

以目前的範例作為背景關係，第一個 `h2` 元素綁定到 12，第二個 `h2` 元素綁定到 31，第三個 `h2` 元素綁定到 22，依此類推。

# --instructions--

修改 `text()` 方法，使每個 `h2` 元素顯示來自 `dataset` 陣列的對應值，並加上一個半形空格及字串 `USD`。例如，第一個標題應該是 `12 USD`。

# --hints--

第一個 `h2` 應該有文字 `12 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

第二個 `h2` 應該有文字 `31 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

第三個 `h2` 應該有文字 `22 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

第四個 `h2` 應該有文字 `17 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

第五個 `h2` 應該有文字 `25 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

第六個 `h2` 應該有文字 `18 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

第七個 `h2` 應該有文字 `29 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

第八個 `h2` 應該有文字 `14 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

第九個 `h2` 應該有文字 `9 USD`。

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
