---
id: 6998b99a8e195cd0c1f211df
title: Passaggio 1
challengeType: 1
dashedName: step-1
---

# --description--

In questo workshop costruirai un semplice tracker per la fauna selvatica usando oggetti JavaScript.

Ecco un esempio di oggetto:

```js
const animal = {}; // empty object
```

Questo crea un oggetto vuoto.

Ora crea una variabile chiamata `tiger` e assegnale un oggetto vuoto.

# --hints--

Dovresti creare una variabile chiamata `tiger`.

```js
assert.exists(tiger);
```

Alla variabile `tiger` dovrebbe essere assegnato un oggetto vuoto.

```js
assert.deepEqual(tiger, {});
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
