---
id: 6999b0d79ea938edcb032237
title: Passaggio 12
challengeType: 1
dashedName: step-12
---

# --description--

In questo ultimo passaggio del workshop Wildlife Tracker, creerai una funzione che consulta una proprietà di un oggetto usando la notazione a parentesi.

Ecco un esempio di come consultare una proprietà usando la notazione a parentesi:

```js
const cat = {
  species: "Cat",
  age: 3
};

const property = "species";

console.log(cat[property]); // "Cat"
```

In questo esempio, la notazione a parentesi ti permette di consultare una proprietà usando una variabile.

Ora crea una funzione chiamata `getProperty`.

La funzione deve avere due parametri: `animal` e `propertyName`.

Restituisci il valore della proprietà usando la notazione a parentesi.

Dopo aver creato la funzione, usa `console.log` per chiamare `getProperty(tiger, "species")` e `getProperty(elephant, "age")` così puoi vedere i valori restituiti nella console.

Con questo, il workshop Wildlife Tracker è completo!

# --hints--

Dovresti creare una funzione chiamata `getProperty`.

```js
assert.isFunction(getProperty);
```

La funzione `getProperty` dovrebbe avere due parametri: `animal` e `propertyName`.

```js
const regex = __helpers.functionRegex('getProperty', ['animal', 'propertyName']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getProperty` dovrebbe usare la notazione a parentesi per consultare la proprietà.

```js
assert.match(code, /animal\s*\[\s*propertyName\s*\]/);
```

Dovresti stampare `getProperty(tiger, "species")` nella console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*tiger\s*,\s*["']species["']\s*\)\s*\)/
);
```

Dovresti stampare `getProperty(elephant, "age")` nella console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*getProperty\s*\(\s*elephant\s*,\s*["']age["']\s*\)\s*\)/
);
```

Chiamare `getProperty(tiger, "species")` dovrebbe restituire `"Tiger"`.

```js
assert.strictEqual(getProperty(tiger, "species"), "Tiger");
```

Chiamare `getProperty(elephant, "age")` dovrebbe restituire `12`.

```js
assert.strictEqual(getProperty(elephant, "age"), 12);
```

`getProperty` dovrebbe usare i parametri della funzione e funzionare con qualsiasi oggetto.

```js
const lion = {
  species: "Lion",
  weight: 200
};

assert.strictEqual(getProperty(lion, "weight"), 200);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

const hasHabitat = (animal) => {
  return animal.hasOwnProperty("habitat");
};

console.log(hasHabitat(tiger));
console.log(hasHabitat(elephant));

const getProperty = (animal, propertyName) => {
  return animal[propertyName];
};

console.log(getProperty(tiger, "species"));
console.log(getProperty(elephant, "age"));
```
