---
id: 69999ff7f7862fdfb885ea5d
title: Passaggio 3
challengeType: 1
dashedName: step-3
---

# --description--

Ora aggiorna il tuo oggetto `tiger` aggiungendo una nuova proprietà chiamata `age`.

Imposta la proprietà `age` al valore `5`.

# --hints--

`tiger` dovrebbe avere una proprietà `age`.

```js
assert.property(tiger, "age");
```

`tiger.age` dovrebbe essere uguale a `5`.

```js
assert.propertyVal(tiger, "age", 5);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
--fcc-editable-region--
  
--fcc-editable-region--
};
```
