---
id: 6999faf886634b72be7a819b
title: Passaggio 2
challengeType: 1
dashedName: step-2
---

# --description--

Per questo tracker per la fauna selvatica, aggiornerai un oggetto con nuove proprietà.

Ecco un esempio di creazione di un oggetto con una proprietà:

```js
const animal = {
  name: "Lion" // property: value
};
```

Ora aggiorna l'oggetto `tiger` in modo che includa una proprietà chiamata `species`.

Imposta il suo valore su `"Tiger"`.

# --hints--

`tiger` dovrebbe essere un oggetto.

```js
assert.isObject(tiger);
```

`tiger` dovrebbe avere una proprietà `species`.

```js
assert.property(tiger, "species");
```

`tiger.species` dovrebbe essere uguale a `"Tiger"`.

```js
assert.propertyVal(tiger, "species", "Tiger");
```

# --seed--

## --seed-contents--

```js
const tiger = {
--fcc-editable-region--
  
--fcc-editable-region--
};
```
