---
id: 6999a0c1698b93d5d41156ed
title: Passaggio 4
challengeType: 1
dashedName: step-4
---

# --description--

Ora aggiungi un'altra proprietà all'oggetto `tiger` chiamata `isEndangered`.

Imposta la proprietà `isEndangered` su `true`.

# --hints--

`tiger` dovrebbe avere una proprietà `isEndangered`.

```js
assert.property(tiger, "isEndangered");
```

`tiger.isEndangered` dovrebbe essere uguale a `true`.

```js
assert.propertyVal(tiger, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
--fcc-editable-region--
  
--fcc-editable-region--
};
```
