---
id: 6999b0198d18d402c36cc5fe
title: Passaggio 10
challengeType: 1
dashedName: step-10
---

# --description--

In questo passaggio, creerai una funzione che rimuove una proprietà da un oggetto.

Ecco un esempio di rimozione di una proprietà usando `delete`:

```js
const bird = {
  species: "Parrot",
  canFly: true
};

const removeFlight = (pet) => {
  delete pet.canFly; // remove property using delete keyword
  return pet; // return the updated object
};

console.log(removeFlight(bird));
// { species: "Parrot" }
```

In questo esempio, la proprietà `canFly` viene rimossa dall'oggetto `bird`.

Ora crea una funzione chiamata `removeEndangeredStatus`. La funzione deve prendere un parametro chiamato `animal`.

All'interno della funzione, rimuovi la proprietà `isEndangered` dall'oggetto `animal` usando la parola chiave `delete`. Ritorna l'oggetto `animal` aggiornato.

Dopo aver creato la funzione, usa `console.log` per chiamare `removeEndangeredStatus(tiger)` così puoi vedere l'oggetto aggiornato nella console.

# --hints--

Dovresti creare una funzione chiamata `removeEndangeredStatus`.

```js
assert.isFunction(removeEndangeredStatus);
```

La funzione `removeEndangeredStatus` dovrebbe avere un solo parametro, `animal`.

```js
const regex = __helpers.functionRegex('removeEndangeredStatus', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

La funzione `removeEndangeredStatus` dovrebbe ritornare l'oggetto `animal` aggiornato.

```js
const testAnimal = { species: "Cat", isEndangered: true };
const result = removeEndangeredStatus(testAnimal);
assert.strictEqual(result, testAnimal);
```

`removeEndangeredStatus` dovrebbe usare la parola chiave `delete` per rimuovere la proprietà `isEndangered`.

```js
assert.match(code, /delete\s+animal\.isEndangered/);
```

Dovresti stampare `removeEndangeredStatus(tiger)` nella console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*removeEndangeredStatus\s*\(\s*tiger\s*\)\s*\)/);
```

Chiamare `removeEndangeredStatus(tiger)` dovrebbe rimuovere la proprietà `isEndangered`.

```js
const updatedTiger = removeEndangeredStatus(tiger);

assert.notProperty(updatedTiger, "isEndangered");
```

`removeEndangeredStatus` dovrebbe usare i parametri della funzione e funzionare con qualsiasi oggetto.

```js
const lion = {
  species: "Lion",
  age: 4,
  isEndangered: true
};

const updatedLion = removeEndangeredStatus(lion);

assert.notProperty(updatedLion, "isEndangered");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

--fcc-editable-region--

--fcc-editable-region--
```
