---
id: 6999b07c34c2cde21ee6e1d8
title: Passaggio 11
challengeType: 1
dashedName: step-11
---

# --description--

In questo passaggio, creerai una funzione che verifica se un oggetto ha una proprietà specifica.

Ecco un esempio di utilizzo di `hasOwnProperty`:

```js
const cat = {
  species: "Cat",
  color: "White"
};

const dog = {
  species: "Dog",
  weight: 50
};

const hasColor = (pet) => {
  return pet.hasOwnProperty("color"); // check if "color" property exists
};

console.log(hasColor(cat)); // true
console.log(hasColor(dog)); // false
```

In questo esempio, la funzione `hasColor` usa `hasOwnProperty` per controllare se la proprietà `color` esiste su un oggetto.

Ora crea una funzione chiamata `hasHabitat`. La funzione deve prendere un parametro chiamato `animal`.

Usa il metodo `hasOwnProperty` per restituire `true` se l'oggetto `animal` ha una proprietà chiamata `"habitat"` e `false` altrimenti.

Dopo aver creato la funzione, usa `console.log` per chiamare `hasHabitat(tiger)` e `hasHabitat(elephant)` così potrai vedere entrambi i risultati nella console.

# --hints--

Dovresti creare una funzione chiamata `hasHabitat`.

```js
assert.isFunction(hasHabitat);
```

La funzione `hasHabitat` dovrebbe avere un solo parametro, `animal`.

```js
const regex = __helpers.functionRegex('hasHabitat', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`hasHabitat` dovrebbe usare il metodo `hasOwnProperty` per controllare la proprietà `habitat`.

```js
assert.match(code, /animal\.hasOwnProperty\s*\(\s*["']habitat["']\s*\)/);
```

Dovresti stampare `hasHabitat(tiger)` nella console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*tiger\s*\)\s*\)/);
```

Dovresti stampare `hasHabitat(elephant)` nella console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*hasHabitat\s*\(\s*elephant\s*\)\s*\)/);
```

Chiamare `hasHabitat(tiger)` dovrebbe restituire `true` dopo aver aggiunto un habitat.

```js
addHabitat(tiger, "Rainforest");
assert.strictEqual(hasHabitat(tiger), true);
```

Chiamare `hasHabitat(elephant)` dovrebbe restituire `false` se non esiste alcun habitat.

```js
const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

assert.strictEqual(hasHabitat(elephant), false);
```

`hasHabitat` dovrebbe usare i parametri della funzione e funzionare con qualsiasi oggetto.

```js
const lion = { species: "Lion", habitat: "Savanna" };
assert.strictEqual(hasHabitat(lion), true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

const updateAge = (animal, newAge) => {
  animal.age = newAge;
  return animal;
};

console.log(updateAge(elephant, 12));

const removeEndangeredStatus = (animal) => {
  delete animal.isEndangered;
  return animal;
};

console.log(removeEndangeredStatus(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
