---
id: 6999ac171928136168654774
title: Passaggio 6
challengeType: 1
dashedName: step-6
---

# --description--

In questo passaggio, creerai una funzione per consultare la proprietà `species` di un oggetto.

Ecco un piccolo esempio di come consultare una proprietà di un oggetto:

```js
const dog = {
  species: "Dog"
};
const getAnimalSpecies = (pet) => {
  return pet.species; // access species using dot notation
};
console.log(getAnimalSpecies(dog)); // Dog
```

In questo esempio, la funzione prende un oggetto come parametro e restituisce la sua proprietà `species`.

Ora crea una funzione chiamata `getSpecies`.

La funzione deve prendere un parametro chiamato `animal`.

Restituisci la proprietà `species` di `animal` usando la notazione a punto.

Dopo aver creato la funzione, usa `console.log` per chiamare `getSpecies(tiger)` così puoi vedere il risultato nella console.

# --hints--

Dovresti creare una funzione chiamata `getSpecies`.

```js
assert.isFunction(getSpecies);
```

La funzione `getSpecies` dovrebbe avere un solo parametro, `animal`.

```js
const regex = __helpers.functionRegex('getSpecies', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

Dovresti stampare `getSpecies(tiger)` nella console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getSpecies\s*\(\s*tiger\s*\)\s*\)/);
```

`getSpecies` dovrebbe consultare la proprietà species usando la notazione a punto.

```js
assert.match(code, /animal\.species/);
```

Chiamare `getSpecies(tiger)` dovrebbe restituire `"Tiger"`.

```js
assert.strictEqual(getSpecies(tiger), "Tiger");
```

`getSpecies` dovrebbe usare il parametro della funzione e funzionare con qualsiasi oggetto.

```js
const lion = { species: "Lion" };
assert.strictEqual(getSpecies(lion), "Lion");
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
