---
id: 6999adc72c02812a28c01e31
title: Passaggio 9
challengeType: 1
dashedName: step-9
---

# --description--

In questo passaggio, creerai una funzione che aggiorna una proprietà di un oggetto.

Ecco un esempio di aggiornamento di una proprietà all'interno di una funzione:

```js
const dog = {
  age: 4
};

const changeAge = (pet, updatedAge) => {
  pet.age = updatedAge; // update existing property using dot notation
  return pet; // return the updated object
}

console.log(changeAge(dog, 6)); // { age: 6 }
```

In questo esempio, la proprietà `age` viene aggiornata con un nuovo valore.

Ora crea una funzione chiamata `updateAge`. La funzione deve prendere due parametri: `animal` e `newAge`.

All'interno della funzione, aggiorna la proprietà `age` dell'oggetto `animal` con `newAge`. Ritorna l'oggetto `animal` aggiornato.

Dopo aver creato la funzione, usa `console.log` per chiamare `updateAge(elephant, 12)` così puoi vedere l'oggetto `elephant` aggiornato nella console.

# --hints--

Dovresti creare una funzione chiamata `updateAge`.

```js
assert.isFunction(updateAge);
```

La funzione `updateAge` dovrebbe avere due parametri: `animal` e `newAge`.

```js
const regex = __helpers.functionRegex('updateAge', ['animal', 'newAge']);
assert.match(__helpers.removeJSComments(code), regex);
```

`updateAge` dovrebbe usare la notazione a punto per aggiornare la proprietà `age`.

```js
assert.match(code, /animal\.age\s*=\s*newAge/);
```

La funzione `updateAge` dovrebbe ritornare l'oggetto `animal` aggiornato.

```js
const testAnimal = { species: "Dog", age: 4 };
const result = updateAge(testAnimal, 6);
assert.strictEqual(result, testAnimal);
```

Dovresti registrare `updateAge(elephant, 12)` nella console.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*updateAge\s*\(\s*elephant\s*,\s*12\s*\)\s*\)/
);
```

Chiamare `updateAge(elephant, 12)` dovrebbe aggiornare la proprietà age.

```js
const updatedElephant = updateAge(elephant, 12);

assert.propertyVal(updatedElephant, "age", 12);
```

`updateAge` dovrebbe usare i parametri della funzione e funzionare con qualsiasi oggetto.

```js
const lion = { species: "Lion", age: 4 };
const updatedLion = updateAge(lion, 9);

assert.propertyVal(updatedLion, "age", 9);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

const getAge = (animal) => {
  return animal.age;
};

console.log(getAge(tiger));

const addHabitat = (animal, habitat) => {
  animal.habitat = habitat;
  return animal;
};

console.log(addHabitat(tiger, "Rainforest"));

--fcc-editable-region--

--fcc-editable-region--
```
