---
id: 699a00de9f564bb0effaa14d
title: Passaggio 7
challengeType: 1
dashedName: step-7
---

# --description--

In questo passaggio, creerai una funzione per consultare la proprietà `age` di un oggetto.

Ecco un esempio:

```js
const cat = {
  age: 3
};

const getAnimalAge = (pet) => {
  return pet.age; // access age using dot notation
}

console.log(getAnimalAge(cat)); // 3
```

In questo esempio, la funzione prende un oggetto come parametro e restituisce la sua proprietà `age`.

Ora crea una funzione chiamata `getAge`.

La funzione deve prendere un parametro chiamato `animal`.

Restituisci la proprietà `age` di `animal` usando la notazione a punto.

Dopo aver creato la funzione, usa `console.log` per chiamare `getAge(tiger)` così puoi vedere il risultato nella console.

# --hints--

Dovresti creare una funzione chiamata `getAge`.

```js
assert.isFunction(getAge);
```

La funzione `getAge` dovrebbe avere un solo parametro, `animal`.

```js
const regex = __helpers.functionRegex('getAge', ['animal']);
assert.match(__helpers.removeJSComments(code), regex);
```

`getAge` dovrebbe consultare la proprietà `age` usando la notazione a punto.

```js
assert.match(code, /animal\.age/);
```

Dovresti stampare `getAge(tiger)` nella console.

```js
assert.match(code, /console\s*\.\s*log\s*\(\s*getAge\s*\(\s*tiger\s*\)\s*\)/);
```

Chiamare `getAge(tiger)` dovrebbe restituire `5`.

```js
assert.strictEqual(getAge(tiger), 5);
```

`getAge` dovrebbe usare il parametro della funzione e funzionare con qualsiasi oggetto.

```js
const lion = { age: 8 };
assert.strictEqual(getAge(lion), 8);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

const elephant = {
  species: "Elephant",
  age: 10,
  isEndangered: true
};

const getSpecies = (animal) => {
  return animal.species;
};

console.log(getSpecies(tiger));

--fcc-editable-region--

--fcc-editable-region--
```
