---
id: 69ba8da6b296391d913096a0
title: Passaggio 5
challengeType: 1
dashedName: step-5
---

# --description--

Ora crea un secondo oggetto chiamato `elephant`.

Aggiungi le seguenti proprietà:

- `species` con il valore `"Elephant"`
- `age` con il valore `10`
- `isEndangered` con il valore `true`

# --hints--

Dovresti creare una variabile chiamata `elephant`.

```js
assert.exists(elephant);
```

`elephant` dovrebbe avere una proprietà `species` uguale a `"Elephant"`.

```js
assert.propertyVal(elephant, "species", "Elephant");
```

`elephant` dovrebbe avere una proprietà `age` uguale a `10`.

```js
assert.propertyVal(elephant, "age", 10);
```

`elephant` dovrebbe avere una proprietà `isEndangered` uguale a `true`.

```js
assert.propertyVal(elephant, "isEndangered", true);
```

# --seed--

## --seed-contents--

```js
const tiger = {
  species: "Tiger",
  age: 5,
  isEndangered: true
};

--fcc-editable-region--

--fcc-editable-region--
```
