---
id: 5f33071498eb2472b87ddee4
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذا المشروع، ستتعلم أساسيات CSS (أوراق الأنماط المتتالية) من خلال بناء قائمة مقهى. CSS هي اللغة المستخدمة لتنسيق مستند HTML. تصف كيف يجب عرض عناصر HTML على الشاشة.

كما تعلمت في الخطوات الأخيرة من تطبيق صور القطط، هناك هيكل أساسي مطلوب لبدء بناء صفحة الويب الخاصة بك. يجب أن يحتوي كل مستند HTML على إعلان `DOCTYPE` وعنصر `html`. يخبر `DOCTYPE` المتصفح بإصدار HTML الذي ينتمي إليه المستند. ويمثل عنصر `html` العنصر الجذري الذي يحتوي على جميع العناصر الأخرى.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

أضف وسم `<!DOCTYPE html>` وعنصر `html` مع خاصية `lang` بقيمة `en`.

# --hints--

يجب أن يكون لديك إعلان `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

يجب أن يكون لديك وسم افتتاحي `<html lang="en">` مع تعيين اللغة إلى الإنجليزية.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يكون لديك وسم إغلاق `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
