---
id: 5f33294a6af5e9188dbdb8f3
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

حان الوقت لإضافة بعض محتوى القائمة. أضف عنصر `main` داخل عنصر `body` الموجود. سيحتوي في النهاية على معلومات التسعير عن القهوة والحلويات التي يقدمها المقهى.

# --hints--

يجب أن تحتوي شفرتك على وسم افتتاحي `<main>`.

```js
assert(code.match(/<main>/i));
```

يجب أن تحتوي شفرتك على وسم إغلاق `</main>`.

```js
assert(code.match(/<\/main>/i));
```

لا تغير عنصر `body`. تأكد من عدم حذف وسم الإغلاق الخاص به عن طريق الخطأ.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

يجب أن يكون وسم `main` داخل وسم `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
