---
id: 5f3313e74582ad9d063e3a38
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف عنصر `head` داخل عنصر `html`، حتى تتمكن من إضافة عنصر `title`. يجب أن يكون نص عنصر `title` هو `Cafe Menu`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<head>`.

```js
assert(code.match(/<head>/i));
```

يجب أن يكون لديك وسم إغلاق `</head>`.

```js
assert(code.match(/<head>/i));
```

يجب أن يكون لديك وسم افتتاحي `<title>`.

```js
assert(code.match(/<title>/i));
```

يجب أن يكون لديك وسم إغلاق `</title>`.

```js
assert(code.match(/<\/title>/i));
```

يجب أن يكون عنصر `title` مسكنًا داخل عنصر `head`.

```js
assert(code.match(/<head>\s*<title>.*<\/title>\s*<\/head>/si));
```

يجب أن يحتوي عنصر `title` على النص `Cafe Menu`. قد تحتاج إلى التحقق من التهجئة.

```js
assert.match(document.querySelector('title')?.innerText, /Cafe Menu/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
