---
id: 5f3326b143638ee1a09ff1e3
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

للاستعداد لإنشاء بعض المحتوى الفعلي، أضف عنصر `body` أسفل عنصر `head`.

# --hints--

يجب أن يكون لديك وسم فتح `<body>`.

```js
assert(code.match(/<body>/i));
```

يجب أن يكون لديك وسم إغلاق `</body>`.

```js
assert(code.match(/<\/body>/i));
```

لا تغير عنصر `head`. تأكد من أنك لم تحذف وسم الإغلاق الخاص به.

```js
assert(code.match(/<head>/i));
assert(code.match(/<\/head>/i));
```

يجب أن يأتي عنصر `body` بعد عنصر `head`.

```js
assert(code.match(/<\/head>[.\n\s]*<body>/im));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```

