---
id: 5f332b23c2045fb843337579
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

لإعلام الزوار بأن المقهى تأسس في 2020، أضف عنصر `p` أسفل عنصر `h1` مع النص `Est. 2020`.

# --hints--

يجب أن يكون لديك وسم فتح `<p>`.

```js
assert(code.match(/<p>/i));
```

يجب أن يكون لديك وسم إغلاق `</p>`.

```js
assert(code.match(/<\/p>/i));
```

لا تغير عنصر `h1` الموجود لديك. تأكد من عدم حذف وسم الإغلاق.

```js
assert.lengthOf(document.querySelectorAll('h1'),1);
```

يجب أن يكون عنصر `p` أسفل عنصر `h1`.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.tagName, 'H1');
```

يجب أن يحتوي عنصر `p` على النص `Est. 2020`.

```js
assert(document.querySelector("p").innerText === "Est. 2020");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
--fcc-editable-region--
      <h1>CAMPER CAFE</h1>
--fcc-editable-region--
    </main>
  </body>
</html>
```

