---
id: 5f331e55dfab7a896e53c3a1
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

يُعد `title` واحدًا من عدة عناصر توفر معلومات إضافية غير مرئية على صفحة الويب، لكنها مفيدة لمحركات البحث أو لكيفية عرض الصفحة.

داخل عنصر `head`، قم بتسكين عنصر `meta` مع خاصية اسمها `charset` بقيمة `utf-8` لإخبار المتصفح بكيفية ترميز الأحرف للصفحة.

تذكر أن عناصر `meta` هي عناصر فارغة لا تحتاج إلى وسم إغلاق.

# --hints--

يجب أن يكون لديك وسم `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

عنصر `meta` هو عنصر فارغ، فلا يجب أن يحتوي على وسم إغلاق `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

يجب أن يحتوي وسم `meta` على خاصية `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

يجب أن تكون قيمة خاصية `charset` هي `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
