---
id: 5f3477ae8466a9a3d2cc953c
title: الخطوة 15
challengeType: 0
dashedName: step-15
---

# --description--

الآن بعد أن أضفت CSS في ملف `styles.css`، احذف عنصر `style` وكل محتوياته. بمجرد حذفه، سيعود النص الذي كان في الوسط إلى جهة اليسار.

# --hints--

يجب ألا يحتوي شفرتك على أي وسم `style`.

```js
assert(!code.match(/style/i));
```

يجب ألا يحتوي ملف HTML على أي محددات CSS.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```

