---
id: 5f344fbc22624a2976425065
title: الخطوة 9
challengeType: 0
dashedName: step-9
---

# --description--

أنشئ عنصر `h2` داخل عنصر `section` وأضف له النص `Coffee`.

# --hints--

يجب أن يكون لديك وسم افتتاحي `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

يجب أن يكون لديك وسم الإغلاق `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

لا تغير عنصر `section` الموجود لديك. تأكد من عدم حذف وسم الإغلاق.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

يجب أن يكون عنصر `h2` داخل عنصر `section` الخاص بك.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

يجب أن يحتوي عنصر `h2` على النص `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
