---
id: 5f3477ae9675db8bb7655b30
title: الخطوة 12
challengeType: 0
dashedName: step-12
---

# --description--

في الخطوة السابقة، استخدمت <dfn>محدِّد النوع</dfn> لتنسيق العنصر `h1`. وسّط محتوى العنصرين `h2` و `p` بإضافة محدِّد نوع جديد لكل منهما إلى عنصر `style` الموجود.

# --hints--

لا تغير محدِّد `h1` الموجود.

```js
const hasH1 = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasH1);
```

لا تضف وسم `style` جديد. أضف قواعد CSS الجديدة إلى وسم `style` الموجود.

```js
assert.isAtMost(document.querySelectorAll('style').length, 2);
```

أضف محدِّد `h2` جديد.

```js
const hasH2 = new __helpers.CSSHelp(document).getStyle('h2');
assert(hasH2);
```

أضف محدِّد `p` جديد.

```js
const hasP = new __helpers.CSSHelp(document).getStyle('p');
assert(hasP);
```

يجب أن يحتوي عنصر `h1` على خاصية `text-align` بقيمة `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

يجب أن يحتوي عنصر `h2` على خاصية `text-align` بقيمة `center`.

```js
const h2TextAlign = new __helpers.CSSHelp(document).getStyle('h2')?.getPropertyValue('text-align');
assert(h2TextAlign === 'center');
```

يجب أن يحتوي عنصر `p` على خاصية `text-align` بقيمة `center`.

```js
const pTextAlign = new __helpers.CSSHelp(document).getStyle('p')?.getPropertyValue('text-align');
assert(pTextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

