---
id: 5f344f9c805cd193c33d829c
title: الخطوة 11
challengeType: 0
dashedName: step-11
---

# --description--

يمكنك إضافة تنسيق إلى عنصر بتحديده في عنصر `style` وضبط خاصية له مثل هذا:

```css
element {
 property: value;
}
```

وسّط محتوى عنصر `h1` بضبط خاصية `text-align` إلى القيمة `center`.

# --hints--

يجب أن يكون لديك محدِّد `h1` في عنصر `style` الخاص بك.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

يجب أن تكون قيمة خاصية `text-align` هي `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

يجب أن يضبط محدِّد `h1` الخاص بك خاصية `text-align` إلى `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

