---
id: 5f356ed6199b0cdef1d2be8f
title: الخطوة 26
challengeType: 0
dashedName: step-26
---

# --description--

حتى الآن كنت تستخدم محددات النوع وid لتنسيق العناصر. مع ذلك، من الشائع أكثر استخدام محدد مختلف لتنسيق العناصر.

يُعرّف <dfn>محدِّد الصنف</dfn> باسم مسبوق بنقطة مباشرة، مثل هذا:

```css
.class-name {
  styles
}
```

غيّر محدد `#menu` الحالي إلى محدد صنف باستبدال `#menu` بصنف اسمه `.menu`.

# --hints--

يجب أن يكون لديك محدد صنف `.menu`.

```js
const hasMenu = new __helpers.CSSHelp(document).getStyle('.menu');
assert(hasMenu);
```

يجب ألا يكون لديك محدد `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle('#menu');
assert(!hasDiv);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

--fcc-editable-region--
#menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
--fcc-editable-region--
```

