---
id: 5f356ed69db0a491745e2bb6
title: الخطوة 27
challengeType: 0
dashedName: step-27
---

# --description--

لتطبيق تنسيق الصنف على عنصر `div`، احذف الخاصية `id` وأضف خاصية `class` إلى وسم البداية لعنصر `div`. تأكد من تعيين قيمة الصنف إلى `menu`.

# --hints--

يجب أن يظل عنصر `div` معروضًا. تأكد من عدم تلف وسم `<div>`.

```js
assert.lengthOf(document.querySelectorAll('div'),1);
```

يجب أن يحتوي عنصر `div` على الصنف `menu`.

```js
assert.isTrue(document.querySelector('div').classList.contains('menu'));
```

يجب ألا يحتوي عنصر `div` الخاص بك على معرف `menu` بعد الآن.

```js
assert.lengthOf(document.querySelectorAll('div#menu'),0);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
--fcc-editable-region--
    <div id="menu">
--fcc-editable-region--
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  /*
  background-color: burlywood;
  */
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```

