---
id: 5f356ed63c7807a4f1e6d054
title: الخطوة 21
challengeType: 0
dashedName: step-21
---

# --description--

الهدف الآن هو جعل العنصر `div` لا يشغل العرض الكامل للصفحة. خاصية CSS `width` مناسبة تمامًا لهذا الغرض.

يمكنك استخدام محدِّد `id` لاستهداف عنصر معين يحمل خاصية `id`. يُعرّف محدِّد `id` بوضع رمز الهاش `#` مباشرة أمام قيمة `id` للعنصر. على سبيل المثال، إذا كان للعنصر خاصية `id` بقيمة `cat`، فستستهدف هذا العنصر هكذا:

```css 
#cat {
  width: 250px;
}
```

استخدم محدِّد `#menu` لإعطاء العنصر عرضًا مقداره `300px`.

# --hints--

يجب أن يكون لديك محدِّد `#menu`.

```js
const hasDiv = new __helpers.CSSHelp(document).getStyle("#menu");
assert(hasDiv);
```

يجب أن تضبط خاصية `width` على القيمة `300px`.

```js
const hasWidth = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style.width === '300px');
assert(hasWidth);
```

يجب أن يكون عرض العنصر `div` الخاص بك 300px.

```js
const divWidth = new __helpers.CSSHelp(document).getStyle("#menu")?.getPropertyValue('width');
assert(divWidth === '300px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div id="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  background-color: burlywood;
}

h1, h2, p {
  text-align: center;
}
--fcc-editable-region--

```

