---
id: 5f3c866de7a5b784048f94b1
title: الخطوة 36
challengeType: 0
dashedName: step-36
---

# --description--

هذا هو ما تريده تقريبًا، لكن سيكون من الأفضل أن يكون النكهة والسعر في نفس السطر. عناصر `p` هي عناصر <dfn>مستوى كتلة</dfn>، لذا تشغل العرض الكامل لعنصر الأصل الخاص بها.

لجعلها في نفس السطر، تحتاج إلى تطبيق بعض التنسيقات على عناصر `p` لتتصرف بشكل أقرب إلى عناصر <dfn>مضمَّن</dfn>.

للقيام بذلك، ابدأ بإضافة خاصية `class` بالقيمة `item` إلى أول عنصر `article` تحت عنوان `Coffee`.

# --hints--

يجب أن تطبق صنف `item` على عنصر `article` الخاص بك.

```js
const el = document.querySelector('article.item');
assert.exists(el);
```

يجب أن يكون لديك عنصر واحد فقط يحمل صنف `item`.

```js
const elements = document.querySelectorAll('.item');
assert.lengthOf(elements, 1);
```

يجب أن يحمل أول عنصر `article` لديك صنف `item`.

```js
assert.equal(document.querySelector('article')?.className, 'item');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
--fcc-editable-region--
          <h2>Coffee</h2>
          <article>
            <p class="flavor">French Vanilla</p>
            <p class="price">3.00</p>
          </article>
--fcc-editable-region--
          <article>
            <p>Caramel Macchiato</p>
            <p>3.75</p>
          </article>
          <article>
            <p>Pumpkin Spice</p>
            <p>3.50</p>
          </article>
          <article>
            <p>Hazelnut</p>
            <p>4.00</p>
          </article>
          <article>
            <p>Mocha</p>
            <p>4.50</p>
          </article>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
body {
  background-image: url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg);
}

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}

.flavor {
  text-align: left;
}

.price {
  text-align: right;
}
```
