---
id: 66cc0de2920bec775f610424
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

Aunque las funciones funcionan como se espera, hay mucha repetición en tu código.

Dado que estás construyendo una calculadora, no deseas tener que crear una función para cada posible combinación de números que quieras sumar.

Sería mejor crear una única función reutilizable que pueda sumar cualquier par de números.

Elimina todo el código que has escrito hasta ahora.

En el siguiente paso, repasarás cómo trabajar con parámetros y argumentos en funciones.

# --hints--

No debes tener una función llamada `addTwoAndSeven`.

```js
assert.notMatch(code, /function\s+addTwoAndSeven\s*\(\s*\)\s*{/);
```

No debes tener una llamada a la función `addTwoAndSeven`.

```js
assert.notMatch(code, /console\.log\(\s*addTwoAndSeven\(\s*\)\s*\)/);
```

No debes tener una función llamada `addThreeAndFour`.

```js
assert.notMatch(code, /function\s+addThreeAndFour\s*\(\s*\)\s*{/);
```

No debes tener una llamada a la función `addThreeAndFour`.

```js
assert.notMatch(code, /console\.log\(\s*addThreeAndFour\(\s*\)\s*\)/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
function addTwoAndSeven() {
  return 2 + 7;
}

console.log(addTwoAndSeven());

function addThreeAndFour() {
  return 3 + 4;
}

console.log(addThreeAndFour());
--fcc-editable-region--
```
