---
id: 66cc2739f687e484d50bb6f1
title: Paso 18
challengeType: 1
dashedName: step-18
---

# --description--

La última pieza de funcionalidad para su calculadora será la funcionalidad de raíz cuadrada.

Un cuadrado es un número multiplicado por sí mismo. Entonces, la raíz cuadrada sería el número que al ser multiplicado por sí mismo da el número original.

Por ejemplo, la raíz cuadrada de `9` es `3` porque `3 * 3 = 9`.

Para obtener la raíz cuadrada de un número en JavaScript, puede usar el método <dfn>Math.sqrt()</dfn>.

Aquí hay un ejemplo:

```js
// result: 3
Math.sqrt(9);
```

Declare una función llamada `calculateSquareRoot` que tome un parámetro `num` y devuelva la raíz cuadrada de `num`.

# --hints--

Debería tener una función llamada `calculateSquareRoot`.

```js
assert.isFunction(calculateSquareRoot);
```

Su función `calculateSquareRoot` debería tomar un parámetro `num`.

```js
const regex = __helpers.functionRegex('calculateSquareRoot', ['num']);
assert.match(__helpers.removeJSComments(code), regex);
```

Su función `calculateSquareRoot` debería devolver la raíz cuadrada de `num`.

```js
assert.equal(calculateSquareRoot(9), 3);
assert.equal(calculateSquareRoot(16), 4);
assert.equal(calculateSquareRoot(25), 5);
```

# --seed--

## --seed-contents--

```js
function calculateSum(num1, num2) {
  return num1 + num2;
}

console.log(calculateSum(2, 5));
console.log(calculateSum(10, 10));
console.log(calculateSum(5, 5));

function calculateDifference(num1, num2) {
  return num1 - num2;
}

console.log(calculateDifference(22, 5));
console.log(calculateDifference(12, 1));
console.log(calculateDifference(17, 9));

function calculateProduct(num1, num2) {
  return num1 * num2;
}

console.log(calculateProduct(13, 5));


function calculateQuotient(num1, num2) {
  return num2 === 0 ? "Error: Division by zero" : num1 / num2;
}

console.log(calculateQuotient(7, 11));
console.log(calculateQuotient(3, 0));

function calculateSquare(num) {
  return num ** 2;
}

console.log(calculateSquare(2));
console.log(calculateSquare(9));

--fcc-editable-region--

--fcc-editable-region--
```
