---
id: 66cc06012e46aa6bc9b8c001
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, revisaremos cómo trabajar con funciones al desarrollar una aplicación calculadora.

En los videos de lección anteriores, aprendiste cómo declarar funciones de esta manera:

```js
// regular function
function myFunction() {

}

// arrow function
const myArrowFunction = () => {

}
```

Comienza creando una función llamada `addTwoAndSeven`.

Puedes elegir entre usar la sintaxis de función regular o la sintaxis de función flecha.

# --hints--

Debes declarar una función llamada `addTwoAndSeven`.

```js
assert.isFunction(addTwoAndSeven);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
