---
id: 587d78b2367417b2b2512b10
title: Elimina elementos usando splice()
challengeType: 1
forumTopicId: 301166
dashedName: remove-items-using-splice
---

# --description--

Bien, ya aprendimos cómo eliminar elementos del principio y del final de los arreglos usando `shift()` y `pop()`, pero ¿qué pasa si queremos eliminar un elemento desde algún punto en medio? ¿O eliminar más de un elemento a la vez? Bueno, ahí es donde entra `splice()`. `splice()` nos permite hacer justo eso: **eliminar cualquier cantidad de elementos consecutivos** desde cualquier parte de un arreglo.

`splice()` puede recibir hasta 3 parámetros, pero por ahora nos enfocaremos solo en los primeros 2. Los dos primeros parámetros de `splice()` son enteros que representan índices, o posiciones, de los elementos en el arreglo sobre el que se llama `splice()`. Y recuerda, los arreglos están *indexados desde cero*, así que para indicar el primer elemento de un arreglo usaríamos `0`. El primer parámetro de `splice()` representa el índice en el arreglo desde donde comenzar a eliminar elementos, mientras que el segundo parámetro indica la cantidad de elementos a eliminar. Por ejemplo:

```js
let array = ['today', 'was', 'not', 'so', 'great'];

array.splice(2, 2);
```

Aquí eliminamos 2 elementos, comenzando con el tercer elemento (en el índice 2). `array` tendría el valor `['today', 'was', 'great']`.

`splice()` no sólo modifica el arreglo que llama, sino que también devuelve un nuevo arreglo que contiene el valor de los elementos eliminados:

```js
let array = ['I', 'am', 'feeling', 'really', 'happy'];

let newArray = array.splice(3, 2);
```

`newArray` tiene el valor `['really', 'happy']`.

# --instructions--

Hemos inicializado un arreglo `arr`. Usa `splice()` para eliminar elementos de `arr`, de forma que sólo contenga elementos que sumen el valor de `10`.

# --hints--

No debes cambiar la línea original de `const arr = [2, 4, 5, 1, 7, 5, 2, 1];`.

```js
assert(
  __helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/constarr=\[2,4,5,1,7,5,2,1\];?/)
);
```

`arr` sólo debe contener elementos que sumen `10`.

```js
assert.strictEqual(
  arr.reduce((a, b) => a + b),
  10
);
```

Tu código debe utilizar el método `splice()` en `arr`.

```js
assert(__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(/));
```

La división (splice) sólo debe eliminar elementos de `arr` y no agregar ningún elemento adicional a `arr`.

```js
assert(
  !__helpers.removeWhiteSpace(__helpers.removeJSComments(code)).match(/arr\.splice\(\d+,\d+,\d+.*\)/g)
);
```

# --seed--

## --seed-contents--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
// Only change code below this line

// Only change code above this line
console.log(arr);
```

# --solutions--

```js
const arr = [2, 4, 5, 1, 7, 5, 2, 1];
arr.splice(1, 4);
```
