---
id: 66f686b8ebdb982fa8e14330
title: Construye un Programa de Conversión de Fechas
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

En este laboratorio, codificarás un programa de conversión de fechas que convierte una fecha dada a diferentes formatos. Por ejemplo, la fecha actual `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)` se formatearía de las siguientes 2 maneras:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**Objetivo:** Cumple las historias de usuario a continuación y supera todas las pruebas para completar el laboratorio.

**Historia de usuario:**

1. Deberías crear una variable llamada `currentDate` y asignarle la fecha y hora actuales usando el objeto `Date`.
2. Deberías crear una variable llamada `currentDateFormat` y asignarle el string `Current Date and Time: [current date]`. Reemplaza `[current date]` con el resultado de la variable `currentDate`.
3. Deberías registrar el valor de `currentDateFormat` en la consola.
4. Deberías crear una función llamada `formatDateMMDDYYYY` que tome el objeto de fecha como un parámetro. Puedes nombrar este parámetro como desees.
5. Tu función `formatDateMMDDYYYY` debería convertir la fecha actual a formato `mes/día/año` y retornar el string `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`.
6. Deberías crear una función llamada `formatDateLong` que tome el objeto de fecha como un parámetro.
7. Tu función `formatDateLong` debería convertir la fecha actual a formato `Mes Día, Año` y retornar el string `Formatted Date (Month Day, Year): [formatted date]`.

**Nota:** Para que las pruebas pasen, asegúrate de usar `en-US` para la localización al formatear las fechas.

# --hints--

Deberías tener una variable llamada `currentDate` que contenga la fecha y hora actuales usando el objeto `Date`.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

Deberías tener una variable llamada `currentDateFormat` que contenga la fecha actual en el formato `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

Deberías registrar el valor de `currentDateFormat` en la consola.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

Deberías tener una función `formatDateMMDDYYYY`

```js
assert.isFunction(formatDateMMDDYYYY);
```

La función `formatDateMMDDYYYY` debería tomar un único parámetro.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

Cuando el objeto de fecha contiene `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la función `formatDateMMDDYYYY` debería retornar `Formatted Date (MM/DD/YYYY): 9/27/2024`.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

Deberías tener una función `formatDateLong`

```js
assert.isFunction(formatDateLong);
```

La función `formatDateLong` debería tomar un único parámetro.

```js
assert.lengthOf(formatDateLong, 1);
```

Cuando el objeto de fecha contiene `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`, la función `formatDateLong` debería retornar `Formatted Date (Month Day, Year): September 27, 2024`.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```

