---
id: 694648abde178bb8202d9511
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Sie sollten nun den Namen der Studentin oder des Studenten im Terminal ausgeben sehen.

Python stellt eine Funktion namens `type()` bereit, mit der Sie den Typ eines Werts überprüfen können.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

Im obigen Beispiel bedeutet die Ausgabe `<class 'str'>`, dass die Variable, die an die `type()`-Funktion übergeben wurde, ein String ist.

Verwenden Sie die `type()`-Funktion mit `name` als Argument und geben Sie die Ausgabe wie im Beispiel aus. Überprüfen Sie die Ausgabe im Terminal, die zeigt, dass `name` vom Typ `str` (String) ist.

# --hints--

Sie sollten das Ergebnis des Aufrufs von `type(name)` ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
