---
id: 694648acde178bb8202d9515
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Fügen Sie nun die Punktzahl des Schülers hinzu.

Deklarieren Sie eine Variable namens `score` und weisen Sie ihr den Wert `80.5` zu.

Obwohl sowohl `age` als auch `score` Zahlen sind, müssen sie nicht vom gleichen Typ sein. Python stellt eine Funktion namens `isinstance()` bereit, um dies zu überprüfen.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Verwenden Sie `isinstance()`, um zu prüfen, ob `score` ein `int` ist, und geben Sie das Ergebnis wie im obigen Beispiel im Terminal aus.

# --hints--

Sie sollten eine `score`-Variable haben.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

Die Variable `score` sollte den Wert `80.5` speichern. Umgeben Sie den Wert nicht mit Anführungszeichen.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Sie sollten das Ergebnis von `isinstance(score, int)` ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
