---
id: 694648acde178bb8202d9514
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Jetzt müssen Sie das Alter des Schülers zum Zeugnis hinzufügen. Dafür verwenden Sie einen Integer, einen der numerischen Datentypen in Python.

Deklarieren Sie eine Variable namens `age` und weisen Sie ihr den Integer-Wert `20` zu.

Geben Sie dann den Wert und den Datentyp von `age` zusammen, durch ein Komma getrennt, aus. Überprüfen Sie die Ausgabe im Terminal, die den Wert von `age` und dessen Typ als `int` (Integer) anzeigt.

# --hints--

Sie sollten eine `age`-Variable haben.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

Die Variable `age` sollte den Wert `20` speichern. Umgeben Sie den Wert nicht mit Anführungszeichen.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Sie sollten `age` und `type(age)` mit einem Komma als Trennzeichen ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
