---
id: 69461e4f98594012939bc1c6
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop üben Sie Python-Datentypen, indem Sie einen einfachen Zeugnisdrucker erstellen.

Wie Sie in vorherigen Lektionen gelernt haben, werden Variablen auf diese Weise zugewiesen:

```py
greeting = 'Hello World'
```

Im obigen Beispiel hat die Variable `greeting` den Wert eines Strings, also einer Folge von Zeichen, die von Anführungszeichen umgeben sind.

Erstellen Sie eine Variable namens `name` und weisen Sie ihr den String `Alice` zu. Denken Sie daran, den Wert wie im Beispiel entweder mit einfachen oder doppelten Anführungszeichen zu umgeben.

**Wichtig:** Python verwendet Einrückungen (die Leerzeichen am Zeilenanfang), um den Code zu strukturieren. Achten Sie in diesem Workshop darauf, dass am Anfang jeder Codezeile keine zusätzlichen Leerzeichen stehen. Zusätzliche Leerzeichen führen zu einem `IndentationError` und verhindern, dass Ihr Code ausgeführt wird.

# --hints--

Sie sollten eine `name`-Variable haben.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

Die Variable `name` sollte den Wert `'Alice'` speichern.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
