---
id: 694648acde178bb8202d9517
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Die Ausgabe ist `True`, was bestätigt, dass `score` ein `float` ist.

Vervollständigen Sie den Zeugnisbericht, indem Sie den `score`-Wert zusammen mit seinem Datentyp mit einer einzigen `print()`-Anweisung ausgeben.

# --hints--

Sie sollten `score` und `type(score)` mit einem Komma als Trennzeichen ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
