---
id: 694648acde178bb8202d9516
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Die Ausgabe ist `False`, was zeigt, dass `score` kein `int` ist.

Eine weitere häufige Art von Zahl in Python ist `float`, die eine Zahl mit Dezimalstellen darstellt. Ersetzen Sie `int` durch `float` im bestehenden `isinstance()`-Aufruf, um dies zu bestätigen.

# --hints--

Sie sollten `int` durch `float` im bestehenden `isinstance(score, int)`-Aufruf ersetzen.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))
--fcc-editable-region--

```
