---
id: 694648acde178bb8202d9513
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Der Name des Lernenden sollte dem gleichen Format wie die anderen Angaben folgen.

Entfernen Sie die vorherigen Ausgaben der Variablen `name`. Geben Sie dann `name` und `type(name)` zusammen in einer Zeile aus, getrennt durch ein Komma, wie im vorherigen Schritt.

# --hints--

Sie sollten `print(name)` nicht in Ihrem Code haben.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

Sie sollten `print(type(name))` nicht in Ihrem Code haben.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Sie sollten `name` und `type(name)` im selben `print()`-Aufruf mit einem Komma als Trennzeichen ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
