---
id: 694648acde178bb8202d9512
title: Schritt 4
challengeType: 20
dashedName: step-4
---

# --description--

Das Zeugnis sollte auch anzeigen, ob der Schüler aktuell eingeschrieben ist. Dies kann mit einem Boolean-Wert dargestellt werden.

Boolean-Werte stellen eine Ja-oder-Nein-Bedingung dar und werden häufig verwendet, um Entscheidungen im Code zu treffen. Es gibt nur zwei Boolean-Werte: `True` und `False`.

Deklarieren Sie eine Variable mit dem Namen `is_student` und weisen Sie ihr den Wert `True` zu.

# --hints--

Sie sollten eine `is_student`-Variable haben.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

Die Variable `is_student` sollte den Wert `True` speichern. Umgeben Sie den Wert nicht mit Anführungszeichen.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
