---
id: 6947a0c2b6de035dd8d3eea1
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

Sie können den Wert einer Variablen mit der `print()`-Funktion ausgeben.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Sie werden in den kommenden Lektionen mehr über Funktionen lernen. Wissen Sie vorerst, dass eine Funktion ein wiederverwendbarer Codeblock ist, der aufgerufen werden kann, um seinen Code auszuführen, und dem Argumente übergeben werden können.

Im obigen Beispiel ist `print(greeting)` ein Funktionsaufruf, und `greeting` ist das Argument der Funktion.

Beziehen Sie sich auf das Beispiel und geben Sie die Variable `name` aus. Überprüfen Sie die Ausgabe im Terminal.

# --hints--

Sie sollten `print()` mit `name` als Argument aufrufen.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
