---
id: 6947a3e516bb8a4a681274a7
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Die `print()`-Funktion kann mehr als einen Wert gleichzeitig anzeigen. Trennen Sie die Werte mit einem Komma (`,`), um sie in derselben Zeile auszugeben.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Geben Sie sowohl `is_student` als auch `type(is_student)` in derselben Zeile mit einem Komma `,` aus, wie im Beispiel gezeigt. Überprüfen Sie dann die Ausgabe im Terminal, die den Wert von `is_student` und dessen Typ als `bool` (Boolean) anzeigt.

# --hints--

Sie sollten `is_student` und `type(is_student)` mit einem Komma als Trennzeichen ausgeben.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
