---
id: 68e6aff5325410118e76e7be
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

在本次工作坊中，你將透過建置可存取的音訊控制元件來練習使用 ARIA。

首先，為你的頁面添加一個 `h1` 元素，內容為 `Audio Controls`。

# --hints--

你應該有一個 `h1` 元素。

```js
assert.exists(document.querySelector('h1'));
```

你的 `h1` 元素應該在你的 `body` 元素內。

```js
assert.exists(document.querySelector('body > h1'));
```

你的 `h1` 元素應該有文字 `Audio Controls`。

```js
const h1Text = document.querySelector('h1')?.innerText.trim();
assert.strictEqual(h1Text, 'Audio Controls');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
