---
id: 68eac57fadcc662ce917e75e
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

接著，新增一個 `button` 元素，並將 `type` 屬性設定為值 `button`。

`button` 元素的文字應該是 `Play`。

# --hints--

你應該有一個 `button` 元素。

```js
assert.exists(document.querySelector('button'));
```

你的 `button` 元素應該有一個 `type` 屬性。

```js
const button = document.querySelector('button');
assert.isTrue(button?.hasAttribute('type'));
```

你的 `button` 元素應該將 `type` 屬性設為 `button`。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.getAttribute('type'), 'button');
```

你的 `button` 元素應該有文字 `Play`。

```js
const button = document.querySelector('button');
assert.strictEqual(button?.innerText.trim(), 'Play');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
