---
id: 68eb7f2397da251a40e04aa2
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

接著，為你的音量控制添加一個 `div` 元素來包含所有元素。

在 `div` 元素內，巢狀一個 `span` 元素，其 `id` 屬性設定為 `volume-label`。此外，`span` 元素的文字應為 `Volume`。

稍後會由 ARIA 屬性取用 `id`，以便輔助技術能將此文字識別為音量控制元件的標籤。

# --hints--

你應該在你的 `button` 元素下方有一個 `div` 元素。

```js
assert.exists(document.querySelector('button + div'));
```

你應該在你的 `div` 元素內有一個 `span` 元素。

```js
assert.exists(document.querySelector('button + div > span'));
```

你的 `span` 元素應該有一個 `id` 屬性。

```js
const spanEl = document.querySelector('button + div > span');
assert.isTrue(spanEl?.hasAttribute('id'));
```

你的 `span` 元素應該有一個 `id` 屬性，其值設定為 `volume-label`。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.getAttribute('id'), 'volume-label');
```

你的 `span` 元素應該有文字 `Volume`。

```js
const spanEl = document.querySelector('button + div > span');
assert.equal(spanEl?.innerText.trim(), 'Volume');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

--fcc-editable-region--

--fcc-editable-region--

  </body>
</html>
```
