---
id: 68ec0a48c9d6db3594d768c3
title: 步驟 4
challengeType: 0
dashedName: step-4
---

# --description--

現在是時候為你的 `span` 元素添加另一個了。

這個 `span` 元素應該設定一個 `id` 屬性為 `volume-description`。此外，這個 `span` 元素的文字應該是 `Adjust the sound level`。

類似先前的 `span` 元素，這個 `span` 的 `id` 稍後會被 ARIA 屬性取用，以便輔助技術能將此文字識別為音量控制的描述。

# --hints--

你應該在 `div` 元素內有兩個 `span` 元素。

```js
const spans = document.querySelectorAll('div > span');
assert.lengthOf(spans, 2);
```

你的第二個 `span` 元素應該有一個 `id` 屬性。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.isTrue(secondSpanEl?.hasAttribute('id'));
```

你的第二個 `span` 元素應該有一個 `id` 屬性，值為 `volume-description`。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.getAttribute('id'), 'volume-description');
```

你的第二個 `span` 元素應該有 `Adjust the sound level` 的文字。

```js
const secondSpanEl = document.querySelector('div > span:nth-of-type(2)');
assert.equal(secondSpanEl?.innerText.trim(), 'Adjust the sound level');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      --fcc-editable-region--
    
      --fcc-editable-region--
    </div>

  </body>
</html>
```
