---
id: 68ed488821845990c911a223
title: 第 7 步
challengeType: 0
dashedName: step-7
---

# --description--

你的可存取音訊控制元件幾乎完成，但還有最後一件事要添加。

在 `div` 元素下方添加一個 `button` 元素，其 `type` 屬性設為 `button`，並且文字為 `Mute`。

有了這個變更，你已完成本次工作坊！

# --hints--

你應該在 `div` 元素下方放置一個 `button` 元素。

```js
const buttonEl = document.querySelector('div + button');
assert.exists(buttonEl);
```

你的 `button` 元素應該有一個 `type` 屬性。

```js
const buttonEl = document.querySelector('div + button');
assert.isTrue(buttonEl?.hasAttribute('type'));
```

你的 `button` 元素應該有一個值為 `button` 的 `type` 屬性。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.getAttribute('type'), 'button');
```

你的 `button` 元素應該有文字 `Mute`。

```js
const buttonEl = document.querySelector('div + button');
assert.equal(buttonEl?.innerText.trim(), 'Mute');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
           aria-labelledby="volume-label volume-description">
    </div>
  --fcc-editable-region--
    
  --fcc-editable-region--

  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
      <input type="range" min="0" max="100" value="50"
             aria-labelledby="volume-label volume-description">
    </div>

    <button type="button">Mute</button>

  </body>
</html>
```
