---
id: 68ed3a96e0ebcc36565b39e5
title: 步驟 5
challengeType: 0
dashedName: step-5
---

# --description--

接著，你需要為你的頁面添加一個 `input` 元素，並將 `type` 屬性設為 `range` 以創建音量滾軸。將 `min` 屬性設為 `0`，`max` 屬性設為 `100`，以及 `value` 屬性設為 `50`，以定義預設音量層級。

# --hints--

你應該在第二個 `span` 元素下方有一個 `input` 元素。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl);
```

你的 `input` 元素應該有一個 `type` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('type'));
```

你的 `input` 元素應該有一個 `type` 屬性，其值為 `range`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('type'), 'range');
```

你的 `input` 元素應該有一個 `min` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.exists(inputEl?.getAttribute('min'));
```

你的 `input` 元素應該有一個值為 `0` 的 `min` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('min'), '0');
```

你的 `input` 元素應該有一個 `max` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('max'));
```

你的 `input` 元素應該有一個值為 `100` 的 `max` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('max'), '100');
```

你的 `input` 元素應該有一個 `value` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('value'));
```

你的 `input` 元素應該有一個 `value` 屬性，且其值為 `50`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.equal(inputEl?.getAttribute('value'), '50');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </div>

  </body>
</html>
```
