---
id: 68ed4204ab7edc66b6cb2383
title: 步驟 6
challengeType: 0
dashedName: step-6
---

# --description--

如你從先前課程中回想，當頁面上已有可用作標籤的文字時，會使用 `aria-labelledby` 屬性。在這種情況下，你已將 `id` 屬性添加到你的 `span` 元素，現在將在此處使用。

為你的 `input` 元素添加一個 `aria-labelledby` 屬性，並設置為 `volume-label volume-description`。

# --hints--

你的 `input` 元素應該有一個 `aria-labelledby` 屬性。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
assert.isTrue(inputEl?.hasAttribute('aria-labelledby'));
```

你的 `input` 元素應該有一個 `aria-labelledby` 屬性，值為 `volume-label volume-description`。

```js
const inputEl = document.querySelector('div > span:nth-of-type(2) + input');
const ariaLabelledbyAttr = inputEl?.getAttribute('aria-labelledby');
assert.oneOf(ariaLabelledbyAttr, ['volume-label volume-description', 'volume-description volume-label']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Accessible Controls</title>
  </head>
  <body>

    <h1>Audio Controls</h1>

    <button type="button">Play</button>

    <div>
      <span id="volume-label">Volume</span>
      <span id="volume-description">Adjust the sound level</span>
    --fcc-editable-region--
    <input type="range" min="0" max="100" value="50">
    --fcc-editable-region--
    </div>

  </body>
</html>
```
