---
id: 587d8259367417b2b2512c85
title: تنفيذ خوارزمية الترتيب بالاختيار
challengeType: 26
dashedName: implement-the-selection-sort-algorithm
---

# --description--

في هذا المختبر ستنفذ خوارزمية الترتيب بالاختيار. تعمل خوارزمية الترتيب بالاختيار عن طريق اختيار القيمة الحد الأدنى في القائمة وتبديلها مع أول قيمة في القائمة. ثم تبدأ من الموضع الثاني، تختار أصغر قيمة في القائمة المتبقية، وتبدلها مع العنصر الثاني. تستمر في التكرار عبر القائمة وتبديل العناصر حتى تصل إلى نهاية القائمة. الآن القائمة مرتبة. خوارزمية الترتيب بالاختيار لها تعقيد زمني تربيعي في جميع الحالات.

**الهدف:** نفذ قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تمتلك دالة `selectionSort`.
1. يجب أن تأخذ دالة `selectionSort` مصفوفة، وتُرجع مصفوفة بنفس العناصر لكن مرتبة من الأصغر إلى الأكبر.

# --hints--

يجب أن تكون `selectionSort` دالة.

```js
assert.isFunction(selectionSort);
```

يجب أن تُرجع `selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92])` مصفوفة لم تتغير إلا في الترتيب.

```js
assert.sameMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]
);
```

يجب أن تُرجع `selectionSort` مصفوفة مرتبة (من الأصغر إلى الأكبر).

```js
assert.sameOrderedMembers(
  selectionSort([1, 4, 2, 8, 345, 123, 43, 32, 5643, 63, 123, 43, 2, 55, 1, 234, 92]),
  [1, 1, 2, 2, 4, 8, 32, 43, 43, 55, 63, 92, 123, 123, 234, 345, 5643]
)
```

يجب ألا تستخدم `selectionSort` الطريقة المدمجة `.sort()`.

```js
function isBuiltInSortUsed(){
  let sortUsed = false;
  const temp = Array.prototype.sort;
  Array.prototype.sort = () => sortUsed = true;
  try {
    selectionSort([0, 1]);
  } finally {
    Array.prototype.sort = temp;
  }
  return sortUsed;
}
assert.isFalse(isBuiltInSortUsed());
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function selectionSort(array) {
  for (let i = 0; i < array.length-1; i++) {
    let minimumIndex = i;
    for (let j = i+1; j < array.length; j++){
      if (array[j] < array[minimumIndex]) {
        minimumIndex = j;
      }
    }
    let value = array[minimumIndex];
    array[minimumIndex] = array[i];
    array[i] = value;
  }
    return array;
}
```
