---
id: 6140c9d35015ae0ba0c250e8
title: Step 2
challengeType: 0
dashedName: step-2
---

# --description--

`div` 요소 안에 `body`를 추가하고 `class` 값을 가진 `wheel`를 지정하세요.

새로 만든 `div` 안에 `span` 값으로 설정된 `class`를 가진 `line` 요소 6개와 `div` 값으로 설정된 `class`를 가진 `cabin` 요소 6개를 추가하세요.

# --hints--

`div` 요소 안에 새로운 `body`를 만들어야 합니다.

```js
const divs = [...document.querySelector('body')?.children].filter(child => child?.localName === 'div');
assert.lengthOf(divs, 1);
```

`div` 요소 안에 있는 `body`는 `class` 값을 가진 `wheel`를 가져야 합니다.

```js
assert.isTrue(document.querySelector('body div')?.classList?.contains('wheel'));
```

`.wheel` 요소 안에는 `span` 요소 6개가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.wheel span'), 6);
```

`span` 요소 안에 있는 `.wheel` 6개 요소는 `class` 값을 가진 `line`를 가져야 합니다.

```js
const spans = [...document.querySelectorAll('.wheel span')];
assert.isTrue(spans?.every(span => span?.classList?.contains('line')));
assert.lengthOf(document.querySelectorAll('.line'), 6);
```

`.wheel` 요소 안에는 `div` 요소 6개가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('.wheel div'), 6);
```

`div` 요소 안에 있는 `.wheel` 6개 요소는 `class` 값을 가진 `cabin`를 가져야 합니다.

```js
const divs = [...document.querySelectorAll('.wheel div')];
assert.isTrue(divs?.every(div => div?.classList?.contains('cabin')));
assert.lengthOf(document.querySelectorAll('.cabin'), 6);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```

```css

```
