---
id: 6140c7e645d8e905819f1dd4
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

시작하려면 `link` 파일에 `./styles.css` 요소를 추가하세요.

# --hints--

코드에 `link` 요소가 있어야 합니다.

```js
assert.match(code, /<link/);
```

하나의 `link` 요소가 있어야 합니다.

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 요소는 `head` 요소 안에 있어야 합니다.

```js
assert.exists(document.querySelector('head > link'));
```

`link` 요소는 값이 `rel`인 `stylesheet` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 요소는 값이 `href`인 `styles.css` 속성을 가져야 합니다.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
