---
id: 6140d10d50636e14695013b2
title: Step 9
challengeType: 0
dashedName: step-9
---

# --description--

두 번째 `.line` 요소를 선택할 선택자를 만드세요. `transform` 속성을 `rotate(60deg)`로 설정하세요.

`transform` 속성은 요소의 모양을 조작할 수 있게 해준다는 점을 기억하세요. 이 경우 `rotate(60deg)` 값을 사용하면 요소가 `transform-origin` 지점을 중심으로 시계 방향으로 60도 회전합니다.

# --hints--

`.line:nth-of-type(2)` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)'));
```

`.line:nth-of-type(2)` 선택자에 `transform` 속성을 `rotate(60deg)`로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(2)')?.transform, 'rotate(60deg)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

--fcc-editable-region--

--fcc-editable-region--
```
