---
id: 6140d0069049f5139d78da40
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

`transform-origin` 속성은 CSS 변형이 적용되는 기준점을 설정하는 데 사용됩니다. 예를 들어, 이 프로젝트에서 나중에 적용할 `rotate` 변형의 경우, `transform-origin`가 요소가 회전하는 기준점을 결정합니다.

`.line` 선택자에 `transform-origin` 속성을 `0% 0%`로 지정하세요. 이렇게 하면 기준점이 왼쪽에서 `0%`, 위쪽에서 `0%`만큼 오프셋되어 요소의 왼쪽 위 모서리로 설정됩니다.

# --hints--

`.line` 선택자에 `transform-origin` 속성을 `0% 0%`로 설정해야 합니다.

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
