---
id: 6140d1a351e88f159ed54fca
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

같은 패턴을 사용하여 세 번째 `.line`, 네 번째 `.line`, 다섯 번째 `.line`, 여섯 번째 `.line`에 대해 각각 별도의 선택자를 만드세요.

세 번째 `transform`의 `.line` 속성을 `rotate(120deg)`으로, 네 번째는 `rotate(180deg)`으로, 다섯 번째는 `rotate(240deg)`으로, 여섯 번째는 `rotate(300deg)`으로 설정하세요.

# --hints--

`.line:nth-of-type(3)` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

`.line:nth-of-type(3)` 선택자에 `transform` 속성을 `rotate(120deg)`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

`.line:nth-of-type(4)` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

`.line:nth-of-type(4)` 선택자에 `transform` 속성을 `rotate(180deg)`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

`.line:nth-of-type(5)` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

`.line:nth-of-type(5)` 선택자에 `transform` 속성을 `rotate(240deg)`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

`.line:nth-of-type(6)` 선택자를 만들어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

`.line:nth-of-type(6)` 선택자에 `transform` 속성을 `rotate(300deg)`으로 설정해야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
