---
id: 6140cdebd39d6a101e747529
title: Step 6
challengeType: 0
dashedName: step-6
---

# --description--

`.line` 요소에 대한 선택자를 만드세요. `background-color`을 `black`로, `width`을 `50%`로, `height`를 `2px`로 설정하는 것부터 시작하세요.

# --hints--

`.line` 선택자가 있어야 합니다.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line'));
```

`.line` 선택자에는 `background-color` 속성이 `black`으로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.backgroundColor, "black");
```

`.line` 선택자에는 `width` 속성이 `50%`으로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.width, "50%");
```

`.line` 선택자에는 `height` 속성이 `2px`으로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line')?.height, "2px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--

--fcc-editable-region--
```
