---
id: 6140f4b5c1555a2960de1e5f
title: Step 22
challengeType: 0
dashedName: step-22
---

# --description--

이름이 `@keyframes`인 또 다른 `cabins` 규칙을 만드세요. `@keyframes wheel`와 동일한 속성을 사용하되, `0%`와 `100%` 규칙을 복사하고 `transform` 선택자의 `100%` 속성을 `rotate(-360deg)`로 설정하세요.

# --hints--

`@keyframes` 규칙이 있어야 합니다.

```js
assert.lengthOf(new __helpers.CSSHelp(document).getCSSRules('keyframes'), 2);
```

새로운 `@keyframes` 규칙의 이름은 `cabins`이어야 합니다.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes');
assert.isTrue(rules?.[0]?.name === 'cabins' || rules?.[1]?.name === 'cabins');
```

새로운 `@keyframes` 규칙은 `@keyframes wheel` 규칙 다음에 와야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.name, 'cabins');
```

`@keyframes wheel` 규칙의 이름은 변경하지 마세요.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[0]?.name, 'wheel');
```

`@keyframes cabins` 규칙에는 `0%` 선택자가 있어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[0]?.keyText, '0%');
```

`0%` 선택자에는 `transform` 속성이 `rotate(0deg)`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[0]?.style?.transform, 'rotate(0deg)');
```

`@keyframes cabins` 규칙에는 `100%` 선택자가 있어야 합니다.

```js
const rules = new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules
assert.isTrue(rules?.[0]?.keyText === '100%' || rules?.[1]?.keyText === '100%');
```

`100%` 선택자는 `0%` 선택자 다음에 와야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[1]?.keyText, '100%')
```

`100%` 선택자에는 `transform` 속성이 `rotate(-360deg)`로 설정되어야 합니다.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getCSSRules('keyframes')?.[1]?.cssRules?.[1]?.style?.transform, 'rotate(-360deg)')
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
  animation-name: wheel;
  animation-duration: 10s;
  animation-iteration-count: infinite;
  animation-timing-function: linear;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(2) {
  right: 17%;
  top: 93.5%;
}
.cabin:nth-of-type(3) {
  right: 67%;
  top: 93.5%;
}
.cabin:nth-of-type(4) {
  left: -8.5%;
  top: 50%;
}
.cabin:nth-of-type(5) {
  left: 17%;
  top: 7%;
}
.cabin:nth-of-type(6) {
  right: 17%;
  top: 7%;
}

@keyframes wheel {
   0% {
     transform: rotate(0deg);
   }
   100% {
     transform: rotate(360deg);
   }
}

--fcc-editable-region--

--fcc-editable-region--
```
