---
id: a5deed1811a43193f9f1c841
title: Implementar um Element Skipper
challengeType: 26
dashedName: implement-an-element-skipper
---

# --description--

Neste laboratório você criará uma função que pula elementos em um array com base em um valor de passo especificado.

**Objetivo**: Cumprir as user stories abaixo e fazer com que todos os testes passem para completar o laboratório.

**Histórias de Usuário**

1. Você deve ter uma função `dropElements` que aceita um array (`arr`) e uma função (`func`) como argumentos.
1. A função `dropElements` deve iterar pelo array e remover elementos começando pelo primeiro até que `func` retorne `true` para um elemento.
1. A função `dropElements` deve retornar os elementos restantes no array se a condição for atendida.
1. Se a condição nunca for satisfeita, deve retornar um array vazio.

# --hints--

Você deve ter uma função `dropElements`.

```js
assert.isFunction(dropElements);
```

`dropElements([1, 2, 3, 4], function(n) {return n >= 3;})` deve retornar `[3, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n >= 3;}), [3, 4]);
```

`dropElements([0, 1, 0, 1], function(n) {return n === 1;})` deve retornar `[1, 0, 1]`.

```js
assert.deepEqual(dropElements([0, 1, 0, 1], function(n) {return n === 1;}), [1, 0, 1]);
```

`dropElements([1, 2, 3], function(n) {return n > 0;})` deve retornar `[1, 2, 3]`.

```js
assert.deepEqual(dropElements([1, 2, 3], function(n) {return n > 0;}), [1, 2, 3]);
```

`dropElements([1, 2, 3, 4], function(n) {return n > 5;})` deve retornar `[]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 4], function(n) {return n > 5;}), []);
```

`dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;})` deve retornar `[7, 4]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 7, 4], function(n) {return n > 3;}), [7, 4]);
```

`dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;})` deve retornar `[3, 9, 2]`.

```js
assert.deepEqual(dropElements([1, 2, 3, 9, 2], function(n) {return n > 2;}), [3, 9, 2]);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function dropElements(arr, func) {
  while (arr.length > 0 && !func(arr[0])) {
    arr.shift();
  }
  return arr;
}
```
