---
id: 61fd7648a7ba2e5882436831
title: Paso 12
challengeType: 0
dashedName: step-12
---

# --description--

Dentro de su elemento `tbody`, agregue cuatro elementos `tr`. Asigne a los primeros tres un atributo `class` establecido en `data`, y al cuarto un atributo `class` establecido en `total`.

# --hints--

Su `tbody` debe tener cuatro elementos `tr`.

```js
const children = [...document.querySelector('tbody')?.children];
assert.lengthOf(children, 4);
children.forEach(child => assert.strictEqual(child?.localName, 'tr'));
```

Los primeros tres elementos `tr` deben tener el atributo `class` establecido en `data`.

```js
const children = [...document.querySelector('tbody')?.children];
children.forEach((child, index) => {
  if (index < 3) {
    assert.isTrue(child?.classList?.contains('data'));
  }
});
```

El cuarto elemento `tr` debe tener el atributo `class` establecido en `total`.

```js
const children = [...document.querySelector('tbody')?.children];
assert.isTrue(children?.[3]?.classList?.contains('total'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
          <table>
            <caption>
              Assets
            </caption>
            <thead>
              <tr>
                <td></td>
                <th><span class="sr-only year">2019</span></th>
                <th><span class="sr-only year">2020</span></th>
                <th class="current"><span class="sr-only year">2021</span></th>
              </tr>
            </thead>
--fcc-editable-region--
            <tbody>
            </tbody>
--fcc-editable-region--
          </table>
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
