---
id: 61fd6ab779390f49148773bb
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Debajo de su elemento `h1`, cree un elemento `div`. Asígnele un atributo `id` establecido en `years`. Desea que este elemento en particular sea oculto de los lectores de pantalla, así que asígnele el atributo `aria-hidden` establecido en `true`.

# --hints--

Debe crear un nuevo elemento `div` después de su elemento `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling?.localName,
  'div'
);
```

Su nuevo elemento `div` debe tener un atributo `id` establecido en `years`.

```js
assert.strictEqual(document.querySelector('div')?.getAttribute('id'), 'years');
```

Su nuevo elemento `div` debe tener el atributo `aria-hidden` establecido en `true`.

```js
assert.strictEqual(
  document.querySelector('div')?.getAttribute('aria-hidden'),
  'true'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
--fcc-editable-region--

--fcc-editable-region--
      </section>
    </main>
  </body>
</html>
```

```css

```
