---
id: 61fd70336ebb3e4f62ee81ba
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Las tablas HTML utilizan el elemento `caption` para describir de qué trata la tabla. El elemento `caption` siempre debe ser el primer hijo de una `table`, pero se puede posicionar con la propiedad CSS `caption-side`.

Agrega un elemento `caption` a tu primer `table` y dale el texto `Assets`.

# --hints--

Tu primer elemento `table` debe tener un elemento `caption`.

```js
assert.strictEqual(
  document.querySelector('table')?.children?.[0]?.localName,
  'caption'
);
```

Tu elemento `caption` debe tener el texto `Assets`.

```js
assert.strictEqual(document.querySelector('caption')?.textContent.trim(), 'Assets');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
          <span class="flex">
            <span>AcmeWidgetCorp</span>
            <span>Balance Sheet</span>
          </span>
        </h1>
        <div id="years" aria-hidden="true">
          <span class="year">2019</span>
          <span class="year">2020</span>
          <span class="year">2021</span>
        </div>
        <div class="table-wrap">
--fcc-editable-region--
          <table>
          </table>
--fcc-editable-region--
          <table></table>
          <table></table>
        </div>
      </section>
    </main>
  </body>
</html>
```

```css

```
