---
id: 61fd67a656743144844941cb
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

Los lectores de pantalla anuncian los elementos HTML basados en el flujo del documento. Eventualmente querremos que el balance tenga un encabezado de "Balance Sheet" y un subencabezado de "AcmeWidgetCorp". Sin embargo, este orden no tiene sentido si es anunciado por un lector de pantalla.

Otorga a tu `span` existente el atributo `class` establecido en `flex`, y agrega dos elementos `span` dentro de él. El primero debe contener el texto `AcmeWidgetCorp`. El segundo debe contener el texto `Balance Sheet`. Usarás CSS para invertir el orden del texto en la página, pero el orden HTML tendrá más sentido para un lector de pantalla.

# --hints--

Tu elemento `span` existente debe tener el atributo `class` configurado como `flex`.

```js
assert.isTrue(
  document.querySelector('h1')?.children?.[0]?.classList?.contains('flex')
);
```

Tu elemento `span` existente debe tener dos nuevos elementos `span` dentro de él.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.localName,
  'span'
);
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.localName,
  'span'
);
```

Tus nuevos elementos `span` no deben tener un atributo `class`.

```js
assert.isEmpty(document.querySelector('.flex')?.children?.[0]?.classList);
assert.isEmpty(document.querySelector('.flex')?.children?.[1]?.classList);
```

Tu primer nuevo elemento `span` debe contener el texto `AcmeWidgetCorp`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[0]?.textContent,
  'AcmeWidgetCorp'
);
```

Tu segundo nuevo elemento `span` debe contener el texto `Balance Sheet`.

```js
assert.strictEqual(
  document.querySelector('.flex')?.children?.[1]?.textContent,
  'Balance Sheet'
);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Balance Sheet</title>
    <link rel="stylesheet" href="./styles.css" />
  </head>
  <body>
    <main>
      <section>
        <h1>
--fcc-editable-region--
          <span>
          </span>
--fcc-editable-region--
        </h1>
      </section>
    </main>
  </body>
</html>
```

```css

```
