---
id: 6133d11ef548f51f876149e3
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

La navegación es una parte fundamental de la accesibilidad, y los lectores de pantalla dependen de ti para proporcionar la estructura de tu página. Esto se logra con elementos HTML semánticos.

Añade un elemento `header` y un `main` a tu página.

El elemento `header` se utilizará para presentar la página, así como para proporcionar un menú de navegación.

El elemento `main` contendrá el contenido principal de tu página.

# --hints--

Deberías agregar un elemento `header` al `body`.

```js
assert.exists(document.querySelector('body > header'));
```

Deberías agregar un elemento `main` al `body`.

```js
assert.exists(document.querySelector('body > main'));
```

El elemento `header` debe estar antes del elemento `main`.

```js
assert.exists(document.querySelector('header + main'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
