---
id: 67d83df6f82eda3868dd2a84
title: بناء شخصية لعبة تقمص أدوار
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

في هذا المختبر ستتمرن على أساسيات Python ببناء تطبيق صغير ينشئ شخصية لمغامرة لعبة تقمص أدوار.

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تملك دالة باسم `create_character`.
1. يجب أن تقبل الدالة، بالترتيب، اسم الشخصية متبوعًا بثلاث معلمات: القوة، الذكاء، والكاريزما.
1. يجب التحقق من صحة اسم الشخصية:
   - إذا لم يكن اسم الشخصية سلسلة نصية، يجب أن ترجع الدالة `The character name should be a string`.
   - إذا كان اسم الشخصية سلسلة نصية فارغة، يجب أن ترجع الدالة `The character should have a name`.
   - إذا كان اسم الشخصية أطول من 10 أحرف، يجب أن ترجع الدالة `The character name is too long`.
   - إذا كان اسم الشخصية يحتوي على مسافات، يجب أن ترجع الدالة `The character name should not contain spaces`.
1. يجب التحقق من صحة المعلمات أيضًا:
   - إذا كانت واحدة أو أكثر من المعلمات ليست أعداد صحيحة، يجب أن ترجع الدالة `All stats should be integers`.
   - إذا كانت واحدة أو أكثر من المعلمات أقل من 1، يجب أن ترجع الدالة `All stats should be no less than 1`.
   - إذا كانت واحدة أو أكثر من المعلمات أكثر من 4، يجب أن ترجع الدالة `All stats should be no more than 4`.
   - إذا لم يكن مجموع كل المعلمات يساوي 7، يجب أن ترجع الدالة `The character should start with 7 points`.
1. إذا اجتازت كل القيم التحقق، يجب أن ترجع الدالة سلسلة نصية مكونة من أربعة أسطر:
   - السطر الأول يجب أن يحتوي على اسم الشخصية.
   - الأسطر من 2 إلى 4 يجب أن تبدأ باختصار المعلمة، `STR`، `INT`، أو `CHA` (بهذا الترتيب)، ثم مسافة، ثم عدد من النقاط الممتلئة (`●`) يساوي قيمة المعلمة، وعدد من النقاط الفارغة (`○`) ليصل المجموع إلى 10. مثال: إذا كانت قيمة القوة 3، يجب أن تكون هناك 3 نقاط ممتلئة تليها 7 نقاط فارغة. النقاط موجودة في المحرر.

إليك السلسلة التي يجب أن تُرجعها الدالة `create_character('ren', 4, 2, 1)`:

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

**ملاحظة:** رغم أن `str` و `int` اختصارات شائعة للمعلمات، يستخدم بايثون هذه الأسماء لأنواع السلاسل النصية والأعداد الصحيحة، لذا من الأفضل تجنب استخدامها كأسماء متغيرات.

# --hints--

يجب أن يكون لديك دالة باسم `create_character`.

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

عندما يُستدعى `create_character` بمعلمة أولى ليست سلسلة نصية، يجب أن ترجع `The character name should be a string`.

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

عندما يُستدعى `create_character` بمعلمة أولى سلسلة نصية، يجب ألا ترجع `The character name should be a string`.

```js
({
    test: () => runPython(
        `assert create_character('ren', 1, 2, 4) != 'The character name should be a string'`
    )
})
```

عند استدعاء `create_character` بمعلمة أولى سلسلة فارغة يجب أن تُرجع `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

عند استدعاء `create_character` بمعلمة أولى ليست سلسلة فارغة يجب ألا تُرجع `The character should have a name`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should have a name'
`
    )
})
```

عندما يُستدعى `create_character` بمعلمة أولى أطول من 10 أحرف، يجب أن ترجع `The character name is too long`.

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

يجب ألا تقول دالة `create_character` أن اسم الشخصية طويل جدًا إذا لم يكن أطول من 10 أحرف.

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

عندما يُستدعى `create_character` بمعلمة أولى تحتوي على مسافة، يجب أن ترجع `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

عندما يُستدعى `create_character` بمعلمة أولى لا تحتوي على مسافة، يجب ألا ترجع `The character name should not contain spaces`.

```js
({
    test: () => runPython(
        `
assert create_character('chacha', 4, 2, 1) != 'The character name should not contain spaces'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية أو الثالثة أو الرابعة التي ليست عددًا صحيحًا، يجب أن ترجع `All stats should be integers`.

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية والثالثة والرابعة وكلها أعداد صحيحة، يجب ألا ترجع `All stats should be integers`.

```js
({
    test: () => runPython(
        `
assert create_character('friend', 4, 2, 1) != 'All stats should be integers'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية أو الثالثة أو الرابعة التي أقل من `1`، يجب أن ترجع `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية والثالثة والرابعة وكلها لا تقل عن `1`، يجب ألا ترجع `All stats should be no less than 1`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no less than 1'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية أو الثالثة أو الرابعة التي تزيد عن `4`، يجب أن ترجع `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية والثالثة والرابعة وكلها لا تزيد عن `4`، يجب ألا ترجع `All stats should be no more than 4`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'All stats should be no more than 4'
`
    )
})
```

عندما لا يكون مجموع المعلمات الثانية والثالثة والرابعة الممررة إلى الدالة `create_character` يساوي `7`، يجب أن ترجع الدالة `create_character` القيمة `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

عندما يُستدعى `create_character` بالمعلمة الثانية والثالثة والرابعة التي مجموعها `7`، يجب ألا ترجع `The character should start with 7 points`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) != 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` يجب أن تُرجع `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`.

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

عندما يُستدعى `create_character` بقيم صحيحة، يجب أن تعرض إحصائيات الشخصية كما هو مطلوب.

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
        return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
