---
id: 66a9c14e3b34c719e34bae20
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ili kuongeza data za mwanafunzi kwenye jedwali, utahitaji kutumia vipengele vya safu ya jedwali na data ya jedwali.

Kipengele cha <dfn>data ya jedwali</dfn>, `td`, kinatumika kuunda seli katika jedwali.

Hapa kuna mfano wa kutumia kipengele cha `td` kwa jedwali la wachezaji wa michezo:

```html
<tr>
  <td>1</td>
  <td>John Doe</td>
  <td>USA</td>
</tr>
```

Ndani ya kipengele chako cha mwili wa jedwali, ongeza kipengele cha safu ya jedwali, `tr`, chenye vipengele vitatu vya data ya jedwali, `td`.

Kipengele cha kwanza cha data ya jedwali kinapaswa kuwa na jina la mwisho la `Davis`.

Kipengele cha pili cha data ya jedwali kinapaswa kuwa na jina la kwanza la `Alex`.

Kipengele cha tatu cha data ya jedwali kinapaswa kuwa na daraja la `54`.

# --hints--

Unapaswa kuwa na kipengele cha `tr` ndani ya kipengele chako cha `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr'));
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi `Davis` ndani ya kipengele cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr td')?.textContent.trim(), 'Davis');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi `Alex` ndani ya kipengele cha `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[1]?.textContent.trim(), 'Alex');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi `54` ndani ya kipengele cha `tr`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr td')[2]?.textContent.trim(), '54');
```

Unapaswa kuwa na vipengele vyote vitatu vya `td` ndani ya kipengele cha `tr`.

```js
assert.match(code, /<tr>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<td>[\s\S]*<\/td>[\s\S]*<\/tr>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
