---
id: 66a9b48a211a73155621d0df
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utazoea kufanya kazi na jedwali za HTML kwa kujenga jedwali la mtihani wa mwisho kwa kundi la wanafunzi.

Katika mafundisho yaliyopita, ulijifunza jinsi ya kufanya kazi na kipengele cha `table` kuwakilisha data ya jedwali.

Ndani ya kipengele chako cha `body`, panga ndani kipengele cha `table`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `table`.

```js
assert.match(code, /<table>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `table`.

```js
assert.match(code, /<\/table>/i);
```

Kipengele chako cha `table` kinapaswa kuwa ndani ya kipengele chako cha `body`.

```js
const tableEl = document.querySelector('table');
assert.strictEqual(tableEl?.parentElement?.tagName, 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
