---
id: 66a9c7a4521d2b1b1ec6dcf0
title: Zoezi la 7
challengeType: 0
dashedName: step-7
---

# --description--

Sasa ni wakati wa kuongeza wanafunzi wawili zaidi kwenye jedwali.

Ukifuata muundo ule ule kama katika hatua iliyopita, ongeza safu ya pili ya jedwali la wanafunzi. Tumia data zifuatazo kwa vipengele vya data vya jedwali:

- Jina la mwisho: `Doe`
- Jina la kwanza: `Samantha`
- Daraja: `92`

Kwa safu ya tatu ya jedwali la mwanafunzi, tumia data zifuatazo kwa vipengele vya data vya jedwali:

- Jina la mwisho: `Rodriguez`
- Jina la kwanza: `Marcus`
- Daraja: `88`

# --hints--

Unapaswa kuwa na `tr` ya pili ndani ya kipengele chako cha `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(2)'));
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Doe` ndani ya `tr` ya pili.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:first-child')?.textContent.trim(), 'Doe');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Samantha` ndani ya `tr` ya pili.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:nth-child(2)')?.textContent.trim(), 'Samantha');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `92` ndani ya `tr` ya pili.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(2) td:last-child')?.textContent.trim(), '92');
```

Unapaswa kuwa na `tr` ya tatu ndani ya kipengele chako cha `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 3);
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Rodriguez` ndani ya `tr` ya tatu.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:first-child')?.textContent.trim(), 'Rodriguez');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Marcus` ndani ya `tr` ya tatu.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:nth-child(2)')?.textContent.trim(), 'Marcus');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `88` ndani ya `tr` ya tatu.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(3) td:last-child')?.textContent.trim(), '88');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
