---
id: 66a9bb8578459a175432e7d0
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Kwa sehemu ya kwanza ya jedwali, utataka kuunganisha yaliyomo ya kichwa yanayowakilisha lebo za safu kwa jina la kwanza la mwanafunzi, jina la mwisho, na alama ya mtihani wa mwisho.

Kipengele cha <dfn>table head</dfn>, `thead`, kinatumika kuunganisha yaliyomo ya kichwa katika jedwali.

Hapa kuna mfano unaotumia kipengele cha `thead`:

```html
<table>
  <thead>
    <!-- header content goes here -->
  </thead>
</table>
```

Chini ya kipengele chako cha `caption`, ongeza kipengele cha kichwa cha jedwali.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `thead`.

```js
assert.match(code, /<thead>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `thead`.

```js
assert.match(code, /<\/thead>/i);
```

Kipengele chako cha `thead` kinapaswa kuwa ndani ya kipengele chako cha `table`.

```js
const theadEl = document.querySelector('thead');
assert.strictEqual(theadEl?.parentElement?.tagName, 'TABLE');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
