---
id: 66a9bcf00f13a418368a272e
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Kipengele cha kichwa cha jedwali kinajumuisha kipengele cha <dfn>safuwima ya jedwali</dfn>, `tr`, ambacho kina vipengele vya <dfn>seli za kichwa cha jedwali</dfn>, `th`.

Hapa kuna mfano unaotumia vipengele vya `tr` na `th` kwa jedwali la michezo:

```html
<table>
  <caption>Football Scores</caption>
  <thead>
    <tr>
      <th>Team</th>
      <th>Wins</th>
      <th>Losses</th>
    </tr>
  </thead>
</table>
```

Ndani ya kipengele chako cha `thead`, ongeza kipengele cha `tr`.

Ndani ya kipengele chako cha `tr`, ongeza vipengele vitatu vya `th`.

Kipengele cha kwanza cha `th` kinapaswa kuwa na maandishi ya `Last Name`. Kipengele cha pili cha `th` kinapaswa kuwa na maandishi ya `First Name`. Kipengele cha tatu cha `th` kinapaswa kuwa na maandishi ya `Grade`.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `tr`.

```js
assert.match(code, /<tr>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `tr`.

```js
assert.match(code, /<\/tr>/i);
```

Kipengele chako cha `tr` kinapaswa kuwa ndani ya kipengele chako cha `thead`.

```js
const trEl = document.querySelector('tr');
assert.strictEqual(trEl?.parentElement?.tagName, 'THEAD');
```

Kipengele chako cha kwanza cha `th` kinapaswa kuwa na maandishi ya `Last Name`.

```js
const thEl = document.querySelectorAll('th')[0];
assert.strictEqual(thEl?.textContent.trim(), 'Last Name');
```

Kipengele chako cha pili cha `th` kinapaswa kuwa na maandishi ya `First Name`.

```js
const thEl = document.querySelectorAll('th')[1];
assert.strictEqual(thEl?.textContent.trim(), 'First Name');
```

Kipengele chako cha tatu cha `th` kinapaswa kuwa na maandishi ya `Grade`.

```js
const thEl = document.querySelectorAll('th')[2];
assert.strictEqual(thEl?.textContent.trim(), 'Grade');
```

Vipengele vyote vya `th` vinapaswa kuwa ndani ya kipengele chako cha `tr`.

```js
assert.match(code, /<tr>.*<th>.*<\/th>.*<th>.*<\/th>.*<th>.*<\/th>.*<\/tr>/is);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
      --fcc-editable-region--
        
      --fcc-editable-region--
      </thead>
    </table>
  </body>
</html>
```
