---
id: 66a9cb26ec6bd41cf6c82bc5
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Sehemu ya mwisho ya kuongeza kwenye jedwali itakuwa kipengele cha <dfn>mguu wa jedwali</dfn>, `tfoot`. Kipengele cha mguu wa jedwali kitumike kuonyesha alama ya wastani kwa wanafunzi wote waliopo kwenye jedwali.

Ongeza kipengele cha `tfoot` kwenye jedwali.

Ndani ya kipengele cha `tfoot`, ongeza kipengele cha `tr`.

Ndani ya kipengele cha `tr`, ongeza vipengele viwili vya `td`.

Kipengele cha kwanza cha `td` kinapaswa kuwa na maandishi ya `Average Grade`. Kipengele cha pili cha `td` kinapaswa kuwa na maandishi ya `78.8`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `tfoot`.

```js
assert.match(code, /<tfoot>/);
```

Unapaswa kuwa na lebo ya kufunga ya `tfoot`.

```js
assert.match(code, /<\/tfoot>/);
```

Unapaswa kuwa na kipengele cha `tr` ndani ya kipengele chako cha `tfoot`.

```js
assert.isNotNull(document.querySelector('tfoot tr'));
```

Kipengele chako cha kwanza cha `td` kinapaswa kuwa na maandishi ya `Average Grade`. Hakikisha kiko ndani ya kipengele cha `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child')?.textContent.trim(), 'Average Grade');
```

Kipengele chako cha pili cha `td` kinapaswa kuwa na maandishi ya `78.8`. Hakikisha kiko ndani ya kipengele cha `tr`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:nth-of-type(2)')?.textContent.trim(), '78.8');
```

Kipengele chako cha `tr` kinapaswa kuwa na vipengele viwili vya `td` ndani yake.

```js
assert.strictEqual(document.querySelectorAll('tfoot tr td').length, 2);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
