---
id: 66a9c006e4f58819396605b3
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sasa umekamilisha sehemu ya kichwa, ni wakati wa kuongeza <dfn>mwili wa jedwali</dfn>, `tbody`. Mwili wa jedwali utawakilisha majina yote ya wanafunzi na alama zao.

Ongeza kipengele cha mwili wa jedwali kwenye jedwali lako.

# --hints--

Unapaswa kuwa na lebo ya ufunguzi ya `tbody`.

```js
assert.match(code, /<tbody>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `tbody`.

```js
assert.match(code, /<\/tbody>/i);
```

Leboa ya ufunguzi ya `tbody` inapaswa kuja kabla ya lebo ya kufunga ya `table`.

```js
assert.match(code, /<tbody>[\s\S]*<\/tbody>/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>
      
--fcc-editable-region--
      
--fcc-editable-region--
    </table>
  </body>
</html>
```
