---
id: 66a9ca099e3ecb1be10a2696
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Sasa ni wakati wa kuongeza wanafunzi wawili zaidi kwenye jedwali.

Ukifuata muundo uleule wa hatua zilizopita, ongeza safu ya nne ya jedwali la wanafunzi. Tumia data zifuatazo kwa vipengele vya data za jedwali:

- Jina la Mwisho: `Thompson`
- Jina la Kwanza: `Jane`
- Daraja: `77`

Kwa safu ya tano ya jedwali la wanafunzi, tumia data zifuatazo kwa vipengele vya data za jedwali:

- Jina la Mwisho: `Williams`
- Jina la Kwanza: `Natalie`
- Daraja: `83`

# --hints--

Unapaswa kuwa na `tr` ya nne ndani ya kipengele chako cha `tbody`.

```js
assert.isNotNull(document.querySelector('tbody tr:nth-of-type(4)'));
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Thompson` ndani ya kipengele cha nne cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:first-child')?.textContent.trim(), 'Thompson');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Jane` ndani ya kipengele cha nne cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(2)')?.textContent.trim(), 'Jane');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `77` ndani ya kipengele cha nne cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(4) td:nth-of-type(3)')?.textContent.trim(), '77');
```

Unapaswa kuwa na `tr` ya tano ndani ya kipengele chako cha `tbody`.

```js
assert.strictEqual(document.querySelectorAll('tbody tr').length, 5);
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Williams` ndani ya kipengele cha tano cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:first-child')?.textContent.trim(), 'Williams');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `Natalie` ndani ya kipengele cha tano cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(2)')?.textContent.trim(), 'Natalie');
```

Unapaswa kuwa na kipengele cha `td` chenye maandishi ya `83` ndani ya kipengele cha tano cha `tr`.

```js
assert.strictEqual(document.querySelector('tbody tr:nth-of-type(5) td:nth-of-type(3)')?.textContent.trim(), '83');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>
        
--fcc-editable-region--
        
--fcc-editable-region--
      </tbody>
    </table>
  </body>
</html>
```
