---
id: 66a9cd4ffa48cf1ef7333640
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Jedwali lako karibu kukamilika lakini kuna jambo moja la mwisho la kuongeza.

Itakuwa nzuri kama kipengele cha `td` kinachotumika kwa `Average Grade` kingeweza kuenea kwa nguzo mbili badala ya moja tu.

Ili kufanya hivyo, unaweza kutumia sifa ya kipengele <dfn>colspan</dfn> kwenye kipengele cha `td`.

Sifa ya kipengele `colspan` inaeleza idadi ya nguzo ambazo seli inapaswa kueneza.

Hapa kuna mfano wa kutumia sifa ya kipengele `colspan` kwa jedwali la michezo:

```html
<tr>
  <td colspan="3">Total Points</td>
</tr>
```

Ndani ya lebo ya kufungua `td`, ongeza sifa ya kipengele `colspan` na uiweke kuwa `"2"`.

Na kwa mabadiliko hayo, jedwali lako limekamilika!

# --hints--

Haupaswi kubadilisha yaliyomo ya kipengele cha `td` kilichopo.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').textContent.trim(), 'Average Grade');
```

Kipengele chako cha `td` kinapaswa kuwa na sifa ya kipengele `colspan`.

```js
assert.isTrue(document.querySelector('tfoot tr td:first-child').hasAttribute('colspan'));
```

Kipengele chako cha `td` kinapaswa kuwa na sifa ya kipengele `colspan` yenye thamani ya `"2"`.

```js
assert.strictEqual(document.querySelector('tfoot tr td:first-child').getAttribute('colspan'), '2');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          --fcc-editable-region--
          <td>Average Grade</td>
          --fcc-editable-region--
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
      <caption>
        Calculus Final Exam Grades
      </caption>

      <thead>     
        <tr>
          <th>Last Name</th>
          <th>First Name</th>
          <th>Grade</th>
        </tr>
      </thead>

      <tbody>
        <tr>
          <td>Davis</td>
          <td>Alex</td>
          <td>54</td>
        </tr>

        <tr>
          <td>Doe</td>
          <td>Samantha</td>
          <td>92</td>
        </tr>

        <tr>
          <td>Rodriguez</td>
          <td>Marcus</td>
          <td>88</td>
        </tr>

        <tr>
          <td>Thompson</td>
          <td>Jane</td>
          <td>77</td>
        </tr>

        <tr>
          <td>Williams</td>
          <td>Natalie</td>
          <td>83</td>
        </tr>
      </tbody>

      <tfoot>
        <tr>
          <td colspan="2">Average Grade</td>
          <td>78.8</td>
        </tr>
      </tfoot>
    </table>
  </body>
</html>
```
