---
id: 66a9b8f14b963916a3baa732
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Ili kuongeza maelezo ya picha kwenye `table`, unaweza kutumia kipengele cha <dfn>maelezo ya picha ya jedwali</dfn>.

Hapa kuna mfano unaotumia kipengele cha `caption`:

```html
<table>
  <caption>Football Scores</caption>
</table>
```

Ndani ya kipengele chako cha `table`, panga ndani kipengele cha `caption` chenye maandishi `Calculus Final Exam Grades`.

# --hints--

Unapaswa kuwa na lebo ya kufungua ya `caption`.

```js
assert.match(code, /<caption>/i);
```

Unapaswa kuwa na lebo ya kufunga ya `caption`.

```js
assert.match(code, /<\/caption>/i);
```

Kipengele chako cha `caption` kinapaswa kuwa ndani ya kipengele chako cha `table`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.parentElement.tagName, 'TABLE');
```

`caption` yako inapaswa kuwa na maandishi `Calculus Final Exam Grades`.

```js
const captionEl = document.querySelector('caption');
assert.strictEqual(captionEl.innerText.trim(), 'Calculus Final Exam Grades');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <title>Calculus Final Exams Table</title>
    <meta charset="UTF-8" />
  </head>
  <body>
    <table>
    --fcc-editable-region--
      
    --fcc-editable-region--      
    </table>
  </body>
</html>
```
