---
id: 67ef8f695db52583424bd118
title: 映画レビューのページをデザインする
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `main` 要素が必要です。
2. `main` 要素の中に、映画のタイトル用の `h1` 要素が必要です。
3. `h1` 要素の下に、映画のカバーを表示する `img` 要素が必要です。`img` 要素には、画像を説明する `alt` テキストが必要です。必要に応じて次の画像を使ってもかまいません: `https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`。
4. 簡単な映画の説明を含む `p` 要素が必要です。
5. 映画の評価を表示する別の `p` 要素が必要です。その中には、以下の項目を順番通りに含めてください:
   - テキスト `strong` を持つ `Movie Rating:` 要素。
   - `span` 属性が `aria-hidden` に設定された `true` 要素で、星を使った評価の視覚的表現 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆` を含みます。
   - その後に、丸括弧内に評価の数値（例: `9.2/10`）が続きます。
6. テキスト `h2` を持つ `Cast Members` 要素が必要です。
7. `ul` 要素が必要です。
8. `ul` 要素の中に、複数の `li` 要素があり、それぞれに俳優名用の `strong` 要素と、テキスト `as` に続く対応するキャラクター名が含まれている必要があります。（例: `James Holloway as Ethan Carter`）

# --hints--

`main` 要素が必要です。

```js
assert.isNotNull(document.querySelector("main"));
```

`main` 要素の中に、映画のタイトルを含む `h1` 要素が必要です。

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

`h1` 要素の下に、説明的な `img` 属性を持つ `alt` 要素が必要です。

```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

画像の後に映画の説明を含む `p` 要素が必要です。

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

評価を含む別の `p` 要素が必要です。

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

2番目の `p` 要素の中に、テキスト `strong` を持つ `Movie Rating: ` 要素が必要です。

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

評価の段落内に `span` 要素が必要です。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

評価の段落内の `span` 要素には、塗りつぶされた星（⭐）または空の星（☆）が10個あり、`span` の後に丸括弧内の数値が続きます。

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

評価の段落内の `span` 要素には、`aria-hidden` に設定された `true` 属性があります。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

`h2` 要素が必要です。

```js
assert.exists(document.querySelector("h2"));
```

2番目の `h2` 要素の後に `p` が必要です。

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

`h2` にはテキスト `Cast Members` が必要です。

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

`ul` の後に `h2` 要素が必要です。

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` には複数の `li` 要素が含まれています。

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

各 `li` には、俳優名用の `strong` 要素と、テキスト `as` に続く対応するキャラクター名が含まれています。（例: `James Holloway as Ethan Carter`）

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
