---
id: 67ef8f695db52583424bd118
title: 映画レビューのページをデザインする
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**目的：** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー：**

1. `main` 要素があるはずです。
2. `main` 要素の中に、映画のタイトル用の `h1` 要素があるはずです。
3. `h1` 要素の下に、映画のカバーを表示する `img` 要素があるはずです。`img` 要素には、画像を説明する `alt` テキストが必要です。お好みで、以下の画像を使用してもかまいません：`https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`。
4. 簡単な映画の説明を含む `p` 要素があるはずです。
5. 映画の評価を表示する別の `p` 要素があるはずです。その中に、記載された順序で次の項目があるはずです：
   - テキストが `Movie Rating:` の `b` 要素。
   - `aria-hidden` 属性が `true` に設定された `span` 要素。星 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆` を使って評価を視覚的に表します。
   - その span の後に、括弧で囲んだ評価を表す数値（例：`9.2/10`）。
6. テキストが `Cast Members` の `h2` 要素があるはずです。
7. `ul` 要素があるはずです。
8. `ul` 要素の中に、複数の `li` 要素があるはずです。各要素には、俳優名を表す `b` 要素があり、その後に `as` というテキストに続けて対応する役名が記載されます。（例：`James Holloway as Ethan Carter`）。

# --hints--

`main` 要素があるはずです。

```js
assert.isNotNull(document.querySelector("main"));
```

`main` 要素の中に、映画のタイトルを含む `h1` 要素が必要です。

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

`h1` 要素の下に、説明的な `alt` 属性を持つ `img` 要素が必要です。
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

画像の後に映画の説明を含む `p` 要素が必要です。

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

評価を含む別の `p` 要素が必要です。

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

2番目の `p` 要素の中に、テキストが `Movie Rating: ` の `b` 要素を用意してください。

```js
const boldEl = document.querySelector("main p:nth-of-type(2) b");
assert.match(boldEl?.innerText.trim(), /^Movie Rating:?$/);
```

評価の段落内に `span` 要素が必要です。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

評価の段落内の `span` 要素には、塗りつぶされた星（⭐）または空の星（☆）が10個あり、`span` の後に丸括弧内の数値が続きます。

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

評価の段落内の `span` 要素には、`aria-hidden` に設定された `true` 属性があります。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

`h2`要素があるはずです。

```js
assert.exists(document.querySelector("h2"));
```

2番目の `h2` 要素の後に `p` が必要です。

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

`h2` にはテキスト `Cast Members` が必要です。

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

`ul` の後に `h2` 要素が必要です。

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` には複数の `li` 要素が含まれています。

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

各 `li` 要素には、俳優名を含む `b` 要素と、その後に `as` というテキストに続いて役名を記述してください。（例：`James Holloway as Ethan Carter`）。

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const bold = li.querySelector("b");
  assert.isNotNull(bold);
  const fullText = li.textContent.trim();
  const actorName = bold.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <b>Movie Rating:</b>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><b>James Holloway</b> as Ethan Carter</li>
        <li><b>Olivia Sterling</b> as Lena Mitchell</li>
        <li><b>William Lancaster</b> as Professor Adams</li>
        <li>
          <b>Maria Collins</b> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
