---
id: 669aff9f5488f1bea056416d
title: 步驟 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In this workshop, you will practice working with semantic HTML by building a blog page dedicated to Mr. Whiskers the cat.

To begin the project, add the `<!DOCTYPE html>`, and an `html` element with a `lang` attribute of `en`.

Remember that you learned how to build a basic HTML boilerplate like this in the previous module.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

# --hints--

你應該有 `<!DOCTYPE html>`。

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

你應該有一個 `html` 開始標籤，語言設置爲英語。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

你應該有一個 `html` 結束標籤。

```js
assert.match(code, /<\/html>/i);
```

你的 `DOCTYPE` 應放在 `html` 元素之前。

```js
assert.match(code, /<!DOCTYPE\s+html>[.\n\s]*<html\s+lang\s*=\s*('|")en\1\s*>/im)
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
