---
id: 669fc938d38e6e38ace9251e
title: 步驟 3
challengeType: 0
dashedName: step-3
---

# --description--

在你的 `head` 元素內，嵌套一個 `meta` 元素，其 `charset` 屬性的值設置爲 `"UTF-8"`。

在該 `meta` 元素下方，添加一個 `title` 元素。

`title` 元素的文本應爲 `Mr. Whiskers' Blog`。

# --hints--

你應該有一個 `meta` 元素。

```js
assert.isNotNull(document.querySelector("meta"));
```

The `meta` element is a void element, it should not have an end tag `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

你的 `meta` 標籤應該有一個 `charset` 屬性。

```js
assert.match(code, /<meta\s+charset\s*/i);
```

你的 `charset` 屬性的值應爲 `"UTF-8"`。

```js
assert.match(code, /charset\s*=\s*('|")UTF-8\1/i);
```

你的 `meta` 元素應嵌套在 `head` 元素中。

```js
const meta = document.querySelector('head > meta');
assert.strictEqual(meta?.parentElement?.tagName, 'HEAD');
```

You should have an opening `title` tag.

```js
assert.match(code, /<title>/i);
```

你應該有一個 `title` 結束標籤。

```js
assert.match(code, /<\/title>/i);
```

Your `title` element should be nested in your `head` element.

```js
assert.match(code, /<head>.*\s*<title>.*<\/title>.*\s*<\/head>/si);
```

你的 `title` 元素應包含文本 `Mr. Whiskers' Blog`。 你可能需要檢查一下拼寫。

```js
const titleText = document.querySelector('title')?.innerText
assert.strictEqual(titleText, "Mr. Whiskers' Blog");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>

  </head>
  --fcc-editable-region--
</html>
```
