---
id: 669fc7e141e4703748c558bf
title: 步驟 2
challengeType: 0
dashedName: step-2
---

# --description--

Inside the `html` element, add a `head` element.

# --hints--

你應該有一個 `head` 開始標籤。

```js
assert.match(code, /<head>/i);
```

You should have a closing `head` tag.

```js
assert.match(code, /<\/head>/i);
```

你的 `head` 開始標籤應在 `head` 結束標籤之前。

```js
assert.match(code, /<head>[.\n\s]*<\/head>/im)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--

--fcc-editable-region--
</html>
```
