---
id: 66629f407d679d3105e8317f
title: Побудуйте гру камінь ножиці папір
challengeType: 25
dashedName: top-build-a-rock-paper-scissors-game
---

# --description--
Ваша гра буде проходити проти комп’ютера. Ви напишете функцію, яка випадковим чином повертає `"rock"`, `"paper"` або `"scissors"`.

Вам не потрібно турбуватися про фронтенд-частину гри. Ви лише напишете логіку гри. Відкрийте `script.js` і починайте слідувати інструкціям.

**Історія користувача:**

1. У вас має бути функція з назвою `getComputerChoice`.

1. Ваша функція `getComputerChoice` має випадковим чином повертати `"rock"`, `"paper"` або `"scissors"`.

**Підказка:** Метод [Math.random](https://developer.mozilla.org/en-US/docs/Web/JavaScript/Reference/Global_Objects/Math/random) повертає випадкове число, яке більше або дорівнює 0 і менше за 1. Подумайте, як ви можете використати це, щоб умовно повертати один із кількох варіантів.

Гру гратиме людина. Ви напишете функцію, яка прийматиме вибір користувача і повертатиме його.

1. Створіть функцію з назвою `getHumanChoice`.

1. Напишіть код так, щоб `getHumanChoice` повертав один із дійсних варіантів залежно від того, що вводить користувач.

**Підказка:** Використовуйте метод [prompt](https://developer.mozilla.org/en-US/docs/Web/API/Window/prompt) для отримання вхідних даних користувача.

Ваша гра відстежуватиме рахунок гравця. Ви напишете змінні для збереження рахунку гравця.

1. Створіть дві нові змінні з назвами `humanScore` і `computerScore` у глобальній області видимості.

1. Ініціалізуйте ці змінні значенням `0`.

Гра проходитиме раунд за раундом. Ви напишете функцію, яка прийматиме вибори людини і комп’ютера як аргументи, зіграє один раунд, збільшить рахунок переможця раунду і виведе оголошення про переможця.

1. Створіть нову функцію з назвою `playRound`.

1. Визначте два параметри для `playRound`. Перший параметр `humanChoice`, другий параметр `computerChoice`. Використайте ці два параметри, щоб приймати вибори людини і комп’ютера як аргументи.

1. Зробіть параметр `humanChoice` вашої функції нечутливим до регістру, щоб гравці могли вводити `"rock"`, `"ROCK"`, `"RocK"` або інші варіації.

1. Напишіть код для вашої функції `playRound`, яка повертає рядок, що представляє переможця раунду.

- Якщо нічия, вона має повертати `"It's a tie!"`.

- Якщо виграє гравець, вона має повертати `"You win! [player choice] beats [computer choice]"`.

- Якщо виграє комп’ютер, вона має повертати `"You lose! [computer choice] beats [player choice]"`.

1. Збільшіть змінну `humanScore` або `computerScore` залежно від переможця раунду.

Гра триватиме 3 раунди. Ви напишете функцію з назвою `playGame`, яка викликає `playRound` для гри у 3 раунди, відстежує рахунки і оголошує переможця наприкінці.

1. Створіть нову функцію з назвою `playGame`.

1. Створіть цикл, який зіграє 3 раунди і кожного разу викликає функцію `playRound` з вибором людини і вибором комп’ютера як аргументами.

1. Наприкінці гри поверніть переможця гри залежно від того, хто виграв більше раундів.

- Якщо людина виграла більше раундів, поверніть повідомлення, що каже `"You win the game!"`.

- Якщо комп’ютер виграв більше раундів, поверніть повідомлення, що каже `"You lose the game!"`.

# --hints--

У вас має бути функція з назвою `getComputerChoice`

```js
assert.isFunction(getComputerChoice);
```

Ваша функція `getComputerChoice` має випадковим чином повертати `"rock"`, `"paper"` або `"scissors"`.

```js
const counts = {}
for (let i = 0; i < 100; i++) {
  const result = getComputerChoice();
  counts[result] = (counts[result] ?? 0) + 1;
}
assert.lengthOf(Object.keys(counts), 3);
assert.include(Object.keys(counts), "rock");
assert.include(Object.keys(counts), "paper"); 
assert.include(Object.keys(counts), "scissors");
```

У вас має бути функція з назвою `getHumanChoice`

```js
assert.isFunction(getHumanChoice);
```

У вас має бути дві змінні з назвами `humanScore` і `computerScore` у глобальній області видимості.

```js
assert.exists(humanScore);
assert.exists(computerScore);
```

У вас має бути функція з назвою `playRound`

```js
assert.isFunction(playRound);
```

Ваша функція `playRound` має приймати вибори людини і комп’ютера як аргументи з параметрами `humanChoice` і `computerChoice`.

```js
assert.match(playRound.toString(), /\s*(?:\bhumanChoice\b\s*,\s*\bcomputerChoice\b)/);
```

Ваша функція `playRound` має бути нечутливою до регістру, щоб гравці могли вводити `"rock"`, `"ROCK"`, `"RocK"` або інші варіації.

```js
assert.match(playRound.toString(), /\bhumanChoice\s*\.toLowerCase\(\)/);
```

Ваша функція `playRound` має повертати рядок `"It's a tie!"`, якщо вибори людини і комп’ютера однакові.

```js
assert.equal(playRound("rock", "rock"), "It's a tie!");
```

Ваша функція `playRound` має повертати рядок `"You win! [player choice] beats [computer choice]"`, якщо виграє гравець.

```js
assert.equal(playRound("rock", "scissors"), "You win! rock beats scissors");
```

Ваша функція `playRound` має повертати рядок `"You lose! [computer choice] beats [player choice]"`, якщо виграє комп’ютер.

```js
assert.equal(playRound("rock", "paper"), "You lose! paper beats rock");
```

Ваша функція `playRound` має збільшувати змінну `humanScore` або `computerScore` залежно від переможця раунду.

```js
humanScore = 0;
computerScore = 0;
playRound("rock", "scissors");
assert.equal(humanScore, 1);
```

У вас має бути функція з назвою `playGame`.

```js
assert.isFunction(playGame);
```

Використайте цикл, щоб зіграти 3 раунди.

```js
assert.match(playGame.toString(), /\bfor\s*\(/);
```

Поверніть переможця гри залежно від того, хто виграв більше раундів.

```js
window.prompt = () => "rock";
assert.match(playGame(), /You (win|lose) the game!/);
```

# --seed--

## --seed-contents--

```js
const hand = ['rock', 'paper', 'scissors'];
```

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

# --solutions--

```html
<html lang="en">
<head>
  <script src="script.js"></script>
</head>
</html>
```

```js
const hand = ['rock', 'paper', 'scissors'];


let computerScore = 0;
let humanScore = 0;

const getComputerChoice = () => {
  return hand[Math.floor(Math.random() * hand.length)];
}

const getHumanChoice = () => {
  return prompt();
}

const playRound = (humanChoice, computerChoice) => {
  humanChoice = humanChoice.toLowerCase()

  const tie = "It's a tie!"
  const win = `You win! ${humanChoice} beats ${computerChoice}`
  const lose = `You lose! ${computerChoice} beats ${humanChoice}`


  if (humanChoice == 'rock') {
    if (computerChoice == 'rock') {
      computerScore++;
      humanScore++;
      return tie;
    } else if (computerChoice == 'paper') {
      computerScore++;
      return lose;
    } else {
      humanScore++;
      return win;
    }
  } else if (humanChoice == 'paper') {
    if (computerChoice == 'rock') {
      humanScore++;
      return win;
    } else if (computerChoice == 'paper') {
      computerScore++;
      humanScore++;
      return tie;
    } else {
      computerScore++;
      return lose;
    }
  } else if (humanChoice == 'scissors') {
    if (computerChoice == 'rock') {
      computerScore++;
      return lose;
    } else if (computerChoice == 'paper') {
      humanScore++;
      return win;
    } else {
      computerScore++;
      humanScore++;
      return tie;
    }
  }
}

const playGame = () => {
  for(let i = 0; i < 3; i++){
    playRound(getHumanChoice(), getComputerChoice())
  }

  return humanScore > computerScore ? "You win the game!" : "You lose the game!"
}

```
