---
id: 69447054012a22f38c944ff0
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Додайте інструкцію `if` всередині вашої функції `countup`.

Перевірте, чи `number` менше за `1`, і поверніть порожній масив.

Це слугуватиме базовим випадком для вашої рекурсії.

# --hints--

Коли `number` менше за `1`, функція має повертати порожній масив.

```js
assert.deepStrictEqual(countup(0), []);
```

Коли `number` **не** менше за `1`, функція **не повинна** нічого повертати.

```js
assert.isUndefined(countup(1)); 
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
