---
id: 69446f18fdcb11aa516f7c5c
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви створите `countup` функцію, яка за допомогою рекурсії створює масив із чисел від 1 до заданого числа.

Почніть зі створення порожньої функції з іменем `countup`, яка приймає один параметр.

# --hints--

У вас має бути функція з іменем `countup`.

```js
assert.isFunction(countup);
```

Ваша функція `countup` має приймати один параметр.

```js
assert.strictEqual(countup.length, 1);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
