---
id: 6948705160b66cf93a916e29
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Оголосіть порожній масив з іменем `countArray` всередині функції `countup` безпосередньо над вашою інструкцією `if`.

Цей масив зберігатиме значення `number` при кожному виклику функції.

# --hints--

Ви повинні оголосити порожній масив з іменем `countArray`.

```js
assert.match(countup.toString(), /(let|const|var)\s+countArray\s*=\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
--fcc-editable-region-- 
  
--fcc-editable-region--
  if (number < 1) {
    return [];
  }
}
```
