---
id: 694870bcf3c6874112b6f951
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Для всіх чисел більших за та рівних `1` створіть блок `else` і всередині нього викличте функцію `countup` з `number - 1` та призначте результат `countArray`. 

Потім використайте `.push()`, щоб додати `number` до `countArray`, і нарешті поверніть `countArray`. 

Це гарантує, що рекурсія врешті-решт досягне базового випадку і завершиться.

# --hints--

Ви повинні призначити результат виклику `countup` з `number - 1` до `countArray`.

```js
assert.match(countup.toString(), /countArray\s*=\s*countup\(\s*number\s*-\s*1\s*\)/);
```

Ви повинні додати `number` до `countArray`.

```js
assert.match(countup.toString(), /countArray\.push\(\s*number\s*\)/);
```

Ви повинні повернути `countArray`.

```js
assert.deepStrictEqual(countup(3), [1, 2, 3]);
```

# --seed--

## --seed-contents--

```js
function countup(number) {
  let countArray = [];
  if (number < 1) {
    return [];
--fcc-editable-region--
  }
  
--fcc-editable-region--
}
```
