---
id: 66f686b8ebdb982fa8e14330
title: 날짜 변환 프로그램 만들기
challengeType: 26
dashedName: lab-date-conversion
---

# --description--

이 실습에서는 주어진 날짜를 다양한 형식으로 변환하는 날짜 변환 프로그램을 코딩할 것입니다. 예를 들어, 현재 날짜 `Fri Sep 27 2024 16:04:43 GMT+0500 (Pakistan Standard Time)`는 다음 두 가지 방식으로 형식화됩니다:

- `(MM/DD/YYYY): 9/27/2024`.
- `(Month Day, Year): September 27, 2024`.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `currentDate` 객체를 사용하여 현재 날짜와 시간을 할당하는 `Date`라는 변수를 만들어야 합니다.
2. `currentDateFormat` 변수의 결과로 `Current Date and Time: [current date]`을 대체하여 문자열 `[current date]`을 할당하는 `currentDate`라는 변수를 만들어야 합니다.
3. `currentDateFormat`의 값을 콘솔에 출력해야 합니다.
4. 날짜 객체를 매개변수로 받는 `formatDateMMDDYYYY`라는 함수를 만들어야 합니다. 이 매개변수 이름은 자유롭게 정할 수 있습니다.
5. `formatDateMMDDYYYY` 함수는 현재 날짜를 `month/day/year` 형식으로 변환하고 문자열 `Formatted Date (MM/DD/YYYY): [month]/[day]/[year]`을 반환해야 합니다.
6. 날짜 객체를 매개변수로 받는 `formatDateLong`라는 함수를 만들어야 합니다.
7. `formatDateLong` 함수는 현재 날짜를 `Month Day, Year` 형식으로 변환하고 문자열 `Formatted Date (Month Day, Year): [formatted date]`을 반환해야 합니다.

**참고:** 테스트가 통과하려면 날짜를 형식화할 때 로케일로 `en-US`를 사용해야 합니다.

# --hints--

`currentDate` 객체를 사용하여 현재 날짜와 시간을 저장하는 `Date`라는 변수가 있어야 합니다.

```js
// excluded timestamps
const expectedDate = new Date();
const expectedDateOnly = new Date(expectedDate.getFullYear(), expectedDate.getMonth(), expectedDate.getDate());

const currentDateOnly = new Date(currentDate.getFullYear(), currentDate.getMonth(), currentDate.getDate());

assert.equal(currentDateOnly.toString(), expectedDateOnly.toString());
```

현재 날짜를 `currentDateFormat` 형식으로 저장하는 `Current Date and Time: <ddd> <MMM> <dd> <yyyy> <HH>:<mm>:<ss> <TIMEZONE>`라는 변수가 있어야 합니다.

```js
const expectedDateString = `Current Date and Time: `;
assert.include(currentDateFormat, expectedDateString);

const currentTimestamp = new Date(currentDateFormat.replace(expectedDateString, '')).getTime();
const expectedTimestamp = new Date().getTime();
assert.approximately(currentTimestamp, expectedTimestamp, 2000);
```

`currentDateFormat`의 값을 콘솔에 출력해야 합니다.

```js
assert.match(code, /console\.log\(\s*currentDateFormat\s*\)/);
```

`formatDateMMDDYYYY` 함수를 만들어야 합니다.

```js
assert.isFunction(formatDateMMDDYYYY);
```

`formatDateMMDDYYYY` 함수는 하나의 매개변수를 받아야 합니다.

```js
assert.lengthOf(formatDateMMDDYYYY, 1);
```

날짜 객체가 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`일 때, `formatDateMMDDYYYY` 함수는 `Formatted Date (MM/DD/YYYY): 9/27/2024`을 반환해야 합니다.

```js
let dateString = "Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)";
let dateObject = new Date(dateString);

assert.equal(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/27/2024");

// prevent hardcoding

dateString = "Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)";
dateObject = new Date(dateString);

assert.strictEqual(formatDateMMDDYYYY(dateObject).toString(), "Formatted Date (MM/DD/YYYY): 9/28/2024");
```

`formatDateLong` 함수를 만들어야 합니다.

```js
assert.isFunction(formatDateLong);
```

`formatDateLong` 함수는 하나의 매개변수를 받아야 합니다.

```js
assert.lengthOf(formatDateLong, 1);
```

날짜 객체가 `Fri Sep 27 2024 16:16:11 GMT+0500 (Pakistan Standard Time)`일 때, `formatDateLong` 함수는 `Formatted Date (Month Day, Year): September 27, 2024`을 반환해야 합니다.

```js
let testDate = new Date('Fri Sep 27 2024 16:21:18 GMT+0500 (Pakistan Standard Time)');
let result = formatDateLong(testDate);

assert.strictEqual(result,'Formatted Date (Month Day, Year): September 27, 2024', `Test failed: ${result}`); 

// prevent hardcoding

testDate = new Date('Fri Sep 28 2024 18:31:18 GMT+0500 (Pakistan Standard Time)');
result = formatDateLong(testDate);
assert.strictEqual(result, 'Formatted Date (Month Day, Year): September 28, 2024'); 
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const currentDate = new Date();
const currentDateFormat = `Current Date and Time: ${currentDate}`; 
console.log(currentDateFormat);

function formatDateMMDDYYYY(date) {
  const month = date.getMonth() + 1;
  const day = date.getDate();
  const year = date.getFullYear();
  return `Formatted Date (MM/DD/YYYY): ${month}/${day}/${year}`;
}

function formatDateLong(date) {
  const longDateFormat =  date.toLocaleDateString('en-US', {
    year: 'numeric',
    month: 'long',
    day: 'numeric'
  });
  return `Formatted Date (Month Day, Year): ${longDateFormat}`;
}

// usage
console.log(formatDateMMDDYYYY(currentDate));
console.log(formatDateLong(currentDate));

```
