---
id: 6140c7e645d8e905819f1dd4
title: Крок 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with the standard boilerplate. Add your `DOCTYPE` declaration, your `html` element with the language set to English, your `head` and `body` elements.

Додайте елемент `meta` для правильного `charset`, елемент `title` та елемент `link` для файлу `./styles.css`.

Встановіть `title` на `Ferris Wheel`.

# --hints--

Код повинен мати оголошення `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE/gi);
```

Додайте пробіл після оголошення `DOCTYPE`.

```js
assert.match(code, /<!DOCTYPE\s+/gi);
```

Визначте тип документа як `html`.

```js
assert.match(code, /<!DOCTYPE\s+html/gi);
```

Закрийте оголошення `DOCTYPE` за допомогою `>` після типу.

```js
assert.match(code, /<!DOCTYPE\s+html\s*>/gi);
```

Елемент `html` повинен мати початковий тег з атрибутом `lang` зі значенням `en`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>/gi);
```

Елемент `html` повинен мати кінцевий тег.

```js
assert.match(code, /<\/html\s*>/);
```

Оголошення `DOCTYPE` має бути на початку HTML.

```js
assert.match(__helpers.removeHtmlComments(code), /^\s*<!DOCTYPE\s+html\s*>/i);
```

Ви повинні мати початковий тег `head`.

```js
assert.match(code, /<head\s*>/i);
```

Ви повинні мати кінцевий тег `head`.

```js
assert.match(code, /<\/head\s*>/i);
```

Ви повинні мати початковий тег `body`.

```js
assert.match(code, /<body\s*>/i);
```

Ви повинні мати кінцевий тег `body`.

```js
assert.match(code, /<\/body\s*>/i);
```

Елемент `body` повинен бути першим елементом після елемента `head`.

```js
assert.match(code, /<\/head\s*>\s*<body\s*>/i);
```

Елементи `head` та `body` повинні бути братськими.

```js
assert.strictEqual(document.querySelector('head')?.nextElementSibling?.localName, 'body');
```

Елемент `head` повинен бути першим елементом в межах елемента `html`.

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>\s*<head\s*>/);
```

Елемент `head` повинен бути в межах елемента `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

Елемент `body` повинен бути в межах елемента `html`.

```js
assert.isTrue([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Код повинен мати елемент `meta`.

```js
const meta = document.querySelector('meta');
assert.exists(meta);
```

Елемент `meta` повинен мати атрибут `charset` зі значенням `UTF-8`.

```js
assert.match(code, /<meta\s+charset=('|"|`)UTF-8\1/i);
```

Елемент `meta` має бути в межах елемента `head`.

```js
const regex = /<head\s*>(?:.|\r|\n)*?<meta\s+charset=('|"|`)utf-8\1(?:.|\r|\n)*?<\/head\s*>/i;
assert.match(code, regex);
```

Your code should have a `title` element.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Your project should have a title of `Ferris Wheel`.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim()?.toLowerCase(), 'ferris wheel');
```

Remember, the casing and spelling matter for the title.

```js
const title = document.querySelector('title');
assert.strictEqual(title?.text?.trim(), 'Ferris Wheel');
```

Your code should have a `link` element.

```js
assert.match(code, /<link/);
```

Ви повинні мати один елемент `link`.

```js
assert.strictEqual(document.querySelectorAll('link').length, 1);
```

Your `link` element should be within your `head` element.

```js
assert.exists(document.querySelector('head > link'));
```

Your `link` element should have a `rel` attribute with the value `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Your `link` element should have an `href` attribute with the value `styles.css`.

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
