---
id: 6735a7370e0ae93a4577c689
title: Diseña la maquetación de una revista
challengeType: 25
dashedName: design-a-magazine-layout
demoType: onClick
---

# --description--

En este laboratorio, diseñarás un esquema de revista usando CSS Grid, incluyendo conceptos como filas y columnas de cuadrícula.

**Objetivo:** Cumple con las historias de usuario a continuación y haz que pasen todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes tener un elemento `main` con la clase `magazine-cover`.

2. Tu `.magazine-cover` debe contener un elemento `header` para el título con la clase `title`.

3. Debes tener cuatro elementos `section` en tu página. Los primeros tres secciones para artículos de revista y un cuarto elemento `section` para una imagen de portada.

4. Los cuatro elementos `section` deben tener un atributo `class` de `feature-article`, `small-article1`, `small-article2`, y `cover-image`, respectivamente.

5. You should set the `display` property of `.magazine-cover` to `grid`.

6. La `.magazine-cover` debe tener una propiedad `grid-template-areas` que define el diseño de la cuadrícula:

   - El título abarca la fila superior.
   - Un artículo destacado que abarca dos columnas y una imagen de portada en la segunda fila.
   - Dos pequeños artículos y la imagen de portada en la tercera fila.

7. La `.magazine-cover` debe tener una propiedad `grid-template-columns` que divide el espacio en tres partes iguales.

8. La `.magazine-cover` debe tener una propiedad `grid-template-rows`. Debe estar configurado con `auto` para la primera fila, y dividir por igual el espacio ocupado por las dos filas restantes.

9. Debes añadir un espacio entre los elementos de la cuadrícula y establecer la anchura de `.magazine-cover` para ocupar la mayor parte del ancho de la ventana gráfica.

**Nota:** Asegúrate de vincular tu hoja de estilos en tu HTML y aplica tu CSS.

# --hints--

Debes tener un elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Tu elemento `main` debe tener la clase `magazine-cover`.

```js
assert.exists(document.querySelector('main.magazine-cover'));
```

Debes tener un elemento `header` dentro de tu elemento `main`.

```js
assert.exists(document.querySelector("main > header"));
```

Tu elemento `header` debe tener la clase `title` y debe contener algún texto.

```js
const header = document.querySelector("header.title");
assert.exists(header);
assert.isAbove(header.innerText.length, 0);
```

Debes tener un elemento `section` para cada artículo y la imagen de portada, con las clases `feature-article`, `small-article1`, `small-article2`, y `cover-image`, respectivamente.

```js
const featureArticle = document.querySelector('section.feature-article');
const smallArticle1 = document.querySelector('section.small-article1');
const smallArticle2 = document.querySelector('section.small-article2');
const coverImage = document.querySelector('section.cover-image');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(coverImage);
```

La propiedad `display` de la `.magazine-cover` debe estar configurada en `grid`.

```js
assert.equal(getComputedStyle(document.querySelector('.magazine-cover')).display, 'grid');
```

Tu `.magazine-cover` debe tener una propiedad `grid-template-areas` con `title` abarcando las tres columnas de la primera fila.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /^("|')title title title\1/);
```

La segunda fila de la plantilla de cuadrícula debe contener un `feature article` que abarque dos columnas y una `cover image`.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /.("|')feature-article feature-article cover-image\1./);
```

La tercera fila de la plantilla de cuadrícula debe contener `small-article1`, `small-article2`, y `cover-image`.

```js
const magazineCoverStyle = getComputedStyle(document.querySelector('.magazine-cover'));
assert.match(magazineCoverStyle.gridTemplateAreas, /("|')small-article1 small-article2 cover-image\1$/);
```

Tu selector `.magazine-cover` debe usar la propiedad `grid-template-columns` para dividir el espacio en tres partes iguales. Recuerda que puedes usar la función `repeat` o definir manualmente cada unidad fraccional.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

Tu clase `.magazine-cover` debe tener una propiedad `grid-template-rows` configurada en `auto 1fr 1fr`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gridTemplateRows, 'auto 1fr 1fr');
```

Debes añadir un espacio de `10px` entre los elementos de la cuadrícula.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.magazine-cover')?.gap, '10px');
```

La clase `.title` debe tener un área de cuadrícula con el mismo nombre.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

La clase `.feature-article` debe tener un área de cuadrícula con el mismo nombre.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

La clase `.cover-image` debe tener un área de cuadrícula con el mismo nombre.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Magazine Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Adventure Magazine</title>
  <link rel="stylesheet" href="styles.css">
</head>
<body>
  <main class="magazine-cover">
      <header class="title">
          <h1>Adventure Awaits</h1>
      </header>
      <section class="feature-article">
          <h2>Top 10 Exotic Destinations for 2024</h2>
          <p>Discover the most breathtaking places to visit this year, from hidden beaches to mountain retreats. Our guide takes you through the best spots for your next adventure.</p>
      </section>
      <section class="small-article1">
          <h3>Gear Up: Must-Have Gadgets</h3>
          <p>Check out the latest tech and gear to make your travels more exciting and comfortable.</p>
      </section>
      <section class="small-article2">
          <h3>Meet the Explorers</h3>
          <p>Get to know the modern adventurers who are pushing the boundaries of exploration.</p>
      </section>
      <section class="cover-image">
          <img src="https://cdn.freecodecamp.org/curriculum/labs/magazine-cover.png" alt="Adventure Cover Image">
      </section>
  </main>
</body>
</html>
```

```css
body {
  font-family: Arial, sans-serif;
  margin: 0;
  padding: 0;
  display: flex;
  justify-content: center;
  align-items: center;
  height: 100vh;
  background-color: #f0f0f0;
}

.magazine-cover {
  display: grid;
  grid-template-areas: 
      "title title title"
      "feature-article feature-article cover-image"
      "small-article1 small-article2 cover-image";
  grid-template-columns: 1fr 1fr 1fr;
  grid-template-rows: auto 1fr 1fr;
  gap: 10px;
  width: 80%;
  background-color: #fff;
}

.title {
  grid-area: title;
  background-color: #1B1B32;
  color: white;
  padding: 10px;
  text-align: center;
}

.feature-article {
  grid-area: feature-article;
  background-color: #f9f9f9;
  padding: 10px;
}

.small-article1, .small-article2 {
  background-color: #e9e9e9;
  padding: 10px;
}

.cover-image {
  grid-area: cover-image;
  display: flex;
  justify-content: center;
  align-items: center;
  background-color: #ddd;
  padding: 10px;
}

.cover-image img {
  max-width: 100%;
  height: auto;
} 
```
