---
id: 691f7bbf76229bb97c60827d
title: Paso 1
challengeType: 1
dashedName: step-1
---

# --description--

En este taller, practicarás el uso de bucles `for...of` creando una función que cuente cuántas veces aparece una cadena en un arreglo de cadenas.

Antes de escribir esa función, crearás una más sencilla que registre cada carácter de una cadena en la consola.

Comienza definiendo una función vacía llamada `printCharacters` con el parámetro `str`.

# --hints--

Debes crear una función `printCharacters`.

```js
assert.isFunction(printCharacters);
```

Tu función `printCharacters` debe tener un parámetro llamado `str`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
