---
id: 6920c9cda1c7f056f75e4b1a
title: Paso 2
challengeType: 1
dashedName: step-2
---

# --description--

Ahora es momento de escribir tu primer ciclo `for...of`. En las lecciones anteriores, aprendiste que puedes usar un ciclo `for...of` para iterar sobre valores de un iterable (por ejemplo, cadenas o arreglos).

Aquí tienes un ejemplo de un ciclo `for...of`:

```js
for (const num of [1, 2, 3]) {
  // code block to be executed
}
```

Agrega un ciclo `for...of` con un bloque de código vacío dentro de tu función. Debe iterar sobre cada carácter del argumento `str`.

# --hints--

Tu función `printCharacters` debe tener un ciclo `for...of` dentro de su bloque de código.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functionprintCharacters\(str\)\{for\(\S+of\S+\)\{\S*?\}\}/);
```

Tu ciclo `for...of` debe iterar sobre el argumento `str` de tu función.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function printCharacters\(str\) \{\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+str/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
