---
id: 6920d189e36b6b60a6330fee
title: Paso 3
challengeType: 1
dashedName: step-3
---

# --description--

Ahora, dentro del ciclo, registra la variable `char` en la consola.

# --hints--

Deberías tener una instrucción `console.log` con `char` como argumento dentro del ciclo.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('printCharacters', ['str'], { capture: true });
const functionCode = codeWithoutJSComments.match(functionRegex)[1];
const normalizedFunctionCode = __helpers.removeWhiteSpace(functionCode);
const loopCodeBlock = normalizedFunctionCode.match(/for\(constcharofstr\)\{(\S*?)}/)[1]

assert.match(loopCodeBlock, /console\.log\(char\)/);
```

Cuando se llama a la función `printCharacters` con el argumento `'test'`, la instrucción `console.log()` dentro de la función debería llamarse cuatro veces con los siguientes argumentos, en orden: `'t'`, `'e'`, `'s'` y `'t'`. En cambio, se llamó con los siguientes argumentos: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
printCharacters('test');
assert.deepEqual(consoleLogSpy.calls, [['t'], ['e'], ['s'], ['t']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
