---
id: 6921d688b8d58581d34459b9
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

En los próximos pasos, crearás una función que cuente cuántas veces aparece una cadena en un arreglo de cadenas.

Para comenzar, define una función vacía llamada `getMatchedWordCount` con los parámetros `sentence` y `match` en ese orden.

# --hints--

Debes crear una función `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Tu función `getMatchedWordCount` debe tener dos parámetros: `sentence` y `match`, en ese orden.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const functionRegex = __helpers.functionRegex('getMatchedWordCount', ['sentence', 'match']);

assert.match(codeWithoutJSComments, functionRegex);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
