---
id: 6921d688b8d58581d34459b9
title: Paso 5
challengeType: 1
dashedName: step-5
---

# --description--

En los próximos pasos, crearás una función que cuente cuántas veces aparece una cadena en un arreglo de cadenas.

Para comenzar, define una función vacía llamada `getMatchedWordCount` con los parámetros `sentence` y `match` en ese orden.

# --hints--

Debes crear una función `getMatchedWordCount`.

```js
assert.isFunction(getMatchedWordCount);
```

Tu función `getMatchedWordCount` debe tener dos parámetros: `sentence` y `match`, en ese orden.

```js
const explorer = await __helpers.Explorer(code);
const { parameters } = explorer.allFunctions.getMatchedWordCount;
assert.lengthOf(parameters, 2);
assert.equal(parameters[0].toString(), 'sentence');
assert.equal(parameters[1].toString(), 'match');
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

--fcc-editable-region--

--fcc-editable-region--
```
