---
id: 6921dae10700e300f26994e3
title: Paso 7
challengeType: 1
dashedName: step-7
---

# --description--

Para ver la salida de tu función `getMatchedWordCount` en su estado actual, añade una instrucción `console.log` debajo. Dentro de `console.log()`, llama a la función con el argumento `["I", "really", "really", "really", "like", "to", "code"]` para el parámetro `sentence` y `"really"` para el parámetro `match`.

# --hints--

Debes llamar a la función `getMatchedWordCount` dentro de una instrucción `console.log`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
assert.match(normalizedCode, /console\.log\(getMatchedWordCount\(\S*?\)\)/);
```

Debes llamar a la función `getMatchedWordCount` con el argumento `["I", "really", "really", "really", "like", "to", "code"]` para el parámetro `sentence` y `"really"` para el parámetro `match`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const normalizedConsoleLogCode = normalizedCode.match(/console\.log\(getMatchedWordCount\(\S*?\)\)/)[0];
assert.match(normalizedConsoleLogCode, /\[('|"|`)I\1,('|"|`)really\2,('|"|`)really\3,('|"|`)really\4,('|"|`)like\5,('|"|`)to\6,('|"|`)code\7\],('|"|`)really\8/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  return count;
}

--fcc-editable-region--

--fcc-editable-region--
```
