---
id: 6921de8d2a111d0341cff702
title: Paso 8
challengeType: 1
dashedName: step-8
---

# --description--

Genial, tu función devuelve un `0`, pero aún no cuenta nada.

Para solucionarlo, dentro de `getMatchedWordCount`, crea un bucle `for...of` con un bloque de código vacío que itere sobre cada palabra en `sentence`.

# --hints--

Tu función `getMatchedWordCount` debería tener un bucle `for...of` dentro de su bloque de código.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);

assert.match(normalizedCode, /functiongetMatchedWordCount\(sentence,match\)\{letcount=0\;for\(\S+of\S+\)\{\S*?\}returncount;\}/);
```

Tu bucle `for...of` debería iterar sobre el argumento `sentence` de tu función.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code); 
const forLoopRegex = /function getMatchedWordCount\(sentence, match\) \{\s*let count = 0\;\s*for\s*\(([\S\s]*?)\)\s*\{[\S\s]*?\}\s*return count\;\s*\}/;
const forLoopCode = codeWithoutJSComments.match(forLoopRegex)[1];

assert.match(forLoopCode, /(const|let|var)\s+\w+\s+of\s+sentence/);
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
--fcc-editable-region--
  
--fcc-editable-region--
  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
