---
id: 6978bd465213832659f3338b
title: Paso 9
challengeType: 1
dashedName: step-9
---

# --description--

Ahora, dentro del ciclo, registra en la consola el siguiente literal de plantilla: `Checking "${word}" against "${match}" | Running count: ${count}`

# --hints--

Debes tener exactamente una declaración `console.log` dentro del ciclo en tu función `getMatchedWordCount`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
console.log(normalizedFunctionCode);

assert.match(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?console\.log\(\S*?\)\S*?\}/);
assert.notMatch(normalizedFunctionCode, /for\(constwordofsentence\)\{\S*?(console\.log\(\S*?\))\S*?(console\.log\(\S*?\))\S*?\}/);
```

Tu declaración `console.log` dentro del ciclo en tu función `getMatchedWordCount` debe usar un literal de plantilla con los siguientes marcadores de posición: `${word}`, `${match}` y `${count}`.

```js
const codeWithoutJSComments = __helpers.removeJSComments(code);
const normalizedCode = __helpers.removeWhiteSpace(codeWithoutJSComments);
const functionRegex = /(functiongetMatchedWordCount\(sentence,match\)\{\S*?returncount;\})/;
const normalizedFunctionCode = normalizedCode.match(functionRegex)[1];
const consoleLogArguments = normalizedFunctionCode.match(/for\(constwordofsentence\)\{\S*?console\.log\((\S*?)\)\S*?\}/)[1]
const templateLiteral = consoleLogArguments.match(/(^\`\S*\`$)/)[1];

assert.include(templateLiteral, "${word}");
assert.include(templateLiteral, "${match}");
assert.include(templateLiteral, "${count}");
```

Cuando se llama a la función `getMatchedWordCount` con `(['foo'], 'bar')`, la declaración `console.log()` dentro de la función debe llamarse una vez con la cadena `Checking "foo" against "bar" | Running count: 0`. En cambio, se llamó con cada cadena en la siguiente lista separada por comas: `--fcc-actual--`.

```js
const consoleLogSpy = __helpers.spyOn(console, 'log');
getMatchedWordCount(['foo'], 'bar');
assert.deepEqual(consoleLogSpy.calls, [['Checking "foo" against "bar" | Running count: 0']]);
consoleLogSpy.restore();
```

# --seed--

## --seed-contents--

```js
function printCharacters(str) {
  for (const char of str) {
    console.log(char);
  }
}
printCharacters("hello");

function getMatchedWordCount(sentence, match) {
  let count = 0;
  
  for (const word of sentence) {
--fcc-editable-region--
    
--fcc-editable-region--
  }

  return count;
}

console.log(
  getMatchedWordCount(
    ["I", "really", "really", "really", "like", "to", "code"],
    "really"
  )
);
```
