---
id: 66f55eac933ff64ce654ca74
title: Construir un Convertidor de Markdown a HTML
challengeType: 25
dashedName: build-a-markdown-to-html-converter
demoType: onClick
---

# --description--

Markdown es un lenguaje de marcado utilizado para agregar elementos de formato a documentos de texto sin formato. Para este laboratorio, todo el HTML y CSS te han sido proporcionados. Utilizarás JavaScript para completar la aplicación Convertidor de Markdown a HTML para que pueda manejar la conversión de construcciones básicas de Markdown en elementos HTML.

**Nota:** El resultado final no será un convertidor de Markdown a HTML integral, pero puedes agregar funcionalidades adicionales si lo deseas.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de Usuario:**

1. Debes tener una función llamada `convertMarkdown` que no tome parámetros.
1. La función `convertMarkdown` debe usar expresiones regulares para convertir la entrada de markdown desde `#markdown-input` en HTML y debe devolver una cadena que contenga el código HTML.
1. La función `convertMarkdown` debe convertir encabezados de nivel uno, dos y tres en los elementos `h1`, `h2` y `h3` correspondientes. Un encabezado en markdown se indica con la cantidad de caracteres `#` que corresponden a su nivel seguido de un espacio y el texto del encabezado. Los caracteres `#` deben colocarse al principio de la línea o estar precedidos por caracteres de espacio.
1. La función `convertMarkdown` debe convertir texto en negrita en elementos `strong`. El texto en negrita en markdown se indica ya sea por un par de asteriscos dobles o un par de guiones bajos dobles que encierran el texto.
1. La función `convertMarkdown` debe convertir texto en cursiva en elementos `em`. El texto en cursiva en markdown se indica ya sea por un par de asteriscos o un par de guiones bajos que encierran el texto.
1. La función `convertMarkdown` debe convertir imágenes en elementos `img`. Una imagen en markdown se indica con `![texto-alt](fuente-imagen)`, donde `texto-alt` es el valor del atributo `alt` y `fuente-imagen` es el valor del atributo `src`.
1. La función `convertMarkdown` debe convertir enlaces en elementos de ancla. Un enlace en markdown se indica con `[texto del enlace](URL)`, donde `texto del enlace` es el texto a encerrar entre las etiquetas de ancla y `URL` es el valor del atributo `href`.
1. La función `convertMarkdown` debe convertir citas en elementos `blockquote`. Una cita en markdown se indica con un `>` seguido por un espacio y el texto de la cita. El carácter `>` debe colocarse al principio de la línea o estar precedido por caracteres de espacio.
1. Cuando introduces texto en `#markdown-input`, el código HTML en bruto devuelto por `convertMarkdown` debe mostrarse dentro de `#html-output`.
1. Cuando introduces texto en `#markdown-input`, el código HTML devuelto por `convertMarkdown` debe renderizarse dentro de `#preview`.

Aquí tienes una tabla que contiene todo el markdown que `convertMarkdown` debería poder manejar y el HTML esperado después de la conversión:

<table>
  <thead>
    <tr>
      <th>Markdown</th>
      <th>HTML</th>
    </tr>
  </thead>
  <tbody>
    <tr>
      <td><code># encabezado 1</code></th>
      <td><code>&lth1&gtencabezado 1&lt/h1&gt</code></th>
    </tr>
    <tr>
      <td><code>## encabezado 2</code></th>
      <td><code>&lth2&gtencabezado 2&lt/h2&gt</code></th>
    </tr>
    <tr>
      <td><code>### encabezado 3</code></th>
      <td><code>&lth3&gtencabezado 3&lt/h3&gt</code></th>
    </tr>
    <tr>
      <td><code>**texto en negrita**</code> o <code>__texto en negrita__</code></th>
      <td><code>&ltstrong&gttexto en negrita&lt/strong&gt</code></th>
    </tr>
    <tr>
      <td><code>*texto en cursiva*</code> o <code>_texto en cursiva_</code></th>
      <td><code>&ltem&gttexto en cursiva&lt/em&gt</code></th>
    </tr>
    <tr>
      <td><code>![alt-texto](fuente-imagen)</code></th>
      <td><code>&ltimg alt="alt-texto" src="fuente-imagen"&gt</code></th>
    </tr>
    <tr>
      <td><code>[texto del enlace](URL)</code></th>
      <td><code>&lta href="URL"&gttexto del enlace&lt/a&gt</code></th>
    </tr>
    <tr>
      <td><code>> cita</code></th>
      <td><code>&ltblockquote&gtcita&lt/blockquote&gt</code></th>
    </tr>
  </tbody>
</table>

**Nota:** Asegúrate de enlazar tu archivo JavaScript en tu HTML.

# --hints--

Debes tener una función llamada `convertMarkdown`.

```js
assert.isFunction(convertMarkdown);
```

Cuando el valor de `#markdown-input` es `# título 1`, `convertMarkdown()` debe devolver `<h1>título 1</h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `# título 1`, `<h1>título 1</h1>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `# título 1`, un elemento `h1` con el texto de `título 1` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# title 1";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h1");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `algún texto # título 1`, `convertMarkdown()` no debe convertir `# título 1` en un elemento `h1`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text # title 1";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH1 = testDiv.querySelector("h1")
assert.notExists(testH1);
```

Cuando el valor de `#markdown-input` es `# título 1` seguido por `# título alternativo` en una nueva línea, `convertMarkdown()` debe devolver `<h1>título 1</h1><h1>título alternativo</h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# title 1\n# alternate title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h1");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 1");
assert.equal(headings[1].innerText, "alternate title");
```

Cuando el valor de `#markdown-input` es `## título 2`, `convertMarkdown()` debe devolver `<h2>título 2</h2>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 2");
```

Cuando el valor de `#markdown-input` es `## título 2`, `<h2>título 2</h2>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

Cuando el valor de `#markdown-input` es `## título 2`, un elemento `h2` con el texto de `título 2` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "## title 2";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h2");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1)
assert.equal(headings[0].innerText, "title 2");
```

Cuando el valor de `#markdown-input` es `algún texto ## título 2`, `convertMarkdown()` no debe convertir `## título 2` en un elemento `h2`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ## title 2";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH2 = testDiv.querySelector("h2")
assert.notExists(testH2);
```

Cuando el valor de `#markdown-input` es `## título 2` seguido por `## título 2 alternativa` en una nueva línea, `convertMarkdown()` debe devolver `<h2>título 2</h2><h2>título 2 alternativa</h2>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "## title 2\n## title 2 alt";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h2");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 2");
assert.equal(headings[1].innerText, "title 2 alt");
```

Cuando el valor de `#markdown-input` es `### título 3`, `convertMarkdown()` debe devolver `<h3>título 3</h3>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Cuando el valor de `#markdown-input` es `### título 3`, `<h3>título 3</h3>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Cuando el valor de `#markdown-input` es `### título 3`, un elemento `h3` con el texto de `título 3` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "### title 3";
input.dispatchEvent(new Event("input"));
const headings = preview.querySelectorAll("h3");
assert.lengthOf(headings, 1);
assert.lengthOf(preview.children, 1);
assert.equal(headings[0].innerText, "title 3");
```

Cuando el valor de `#markdown-input` es `algún texto ### título 3`, `convertMarkdown()` no debe convertir `### título 3` en un elemento `h3`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text ### title 3";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testH3 = testDiv.querySelector("h3")
assert.notExists(testH3);
```

Cuando el valor de `#markdown-input` es `### título 3` seguido por `### tercer título` en una nueva línea, `convertMarkdown()` debe devolver `<h3>título 3</h3><h3>tercer título</h3>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "### title 3\n### third title";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const headings = testDiv.querySelectorAll("h3");
assert.lengthOf(headings, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(headings[0].innerText, "title 3");
assert.equal(headings[1].innerText, "third title");
```

Cuando el valor de `#markdown-input` es `**esto es negrita**`, `convertMarkdown()` debe devolver `<strong>esto es negrita</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(strongs[0].innerText, "this is bold");
```

Cuando el valor de `#markdown-input` es `**esto es negrita**`, `<strong>esto es negrita</strong>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Cuando el valor de `#markdown-input` es `**esto es negrita**`, un elemento `strong` con el texto de `esto es negrita` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "**this is bold**";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Cuando el valor de `#markdown-input` es `**esto es negrita**` seguido por `**esto también es negrita**` en una nueva línea, `convertMarkdown()` debe devolver `<strong>esto es negrita</strong><strong>esto también es negrita</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "**this is bold**\n**this is also bold**";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

Cuando el valor de `#markdown-input` es `__esto es negrita__`, `<strong>esto es negrita</strong>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(testDiv.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Cuando el valor de `#markdown-input` es `__esto es negrita__`, un elemento `strong` con el texto de `esto es negrita` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "__this is bold__";
input.dispatchEvent(new Event("input"));
const strongs = preview.querySelectorAll("strong");
assert.lengthOf(strongs, 1);
assert.lengthOf(preview.children, 1)
assert.equal(strongs[0].innerText, "this is bold");
```

Cuando el valor de `#markdown-input` es `__esto es negrita__` seguido por `__esto también es negrita__` en una nueva línea, `convertMarkdown()` debe devolver `<strong>esto es negrita</strong><strong>esto también es negrita</strong>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "__this is bold__\n__this is also bold__";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const strongs = testDiv.querySelectorAll("strong");
assert.lengthOf(strongs, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(strongs[0].innerText, "this is bold");
assert.equal(strongs[1].innerText, "this is also bold");
```

Cuando el valor de `#markdown-input` es `*esto es cursiva*`, `convertMarkdown()` debe devolver `<em>esto es cursiva</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `*esto es cursiva*`, `<em>esto es cursiva</em>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `*esto es cursiva*`, un elemento `em` con el texto de `esto es cursiva` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "*this is italic*";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `*esto es cursiva*` seguido por `*esto también es cursiva*` en una nueva línea, `convertMarkdown()` debe devolver `<em>esto es cursiva</em><em>esto también es cursiva</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "*this is italic*\n*this is also italic*";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

Cuando el valor de `#markdown-input` es `_esto es cursiva_`, `convertMarkdown()` debe devolver `<em>esto es cursiva</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `_esto es cursiva_`, `<em>esto es cursiva</em>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(testDiv.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `_esto es cursiva_`, un elemento `em` con el texto de `esto es cursiva` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "_this is italic_";
input.dispatchEvent(new Event("input"));
const italics = preview.querySelectorAll("em");
assert.lengthOf(italics, 1);
assert.lengthOf(preview.children, 1);
assert.equal(italics[0].innerText, "this is italic");
```

Cuando el valor de `#markdown-input` es `_esto es cursiva_` seguido por `_esto también es cursiva_` en una nueva línea, `convertMarkdown()` debe devolver `<em>esto es cursiva</em><em>esto también es cursiva</em>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "_this is italic_\n_this is also italic_";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const italics = testDiv.querySelectorAll("em");
assert.lengthOf(italics, 2);
assert.lengthOf(testDiv.children, 2);
assert.equal(italics[0].innerText, "this is italic");
assert.equal(italics[1].innerText, "this is also italic");
```

Cuando el valor de `#markdown-input` es `# **título 1**` o `# __título 1__`, `convertMarkdown()` debe devolver `<h1><strong>título 1</strong></h1>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "# **title 1**";
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = convertMarkdown();
let headings = testDiv1.querySelectorAll("h1");
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

input.value = "# __title 1__";
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = convertMarkdown();
headings = testDiv2.querySelectorAll("h1");
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `# **título 1**` o `# __título 1__`, `<h1><strong>título 1</strong></h1>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
const testDiv1 = document.createElement("div");
testDiv1.innerHTML = output.innerText;
let headings = testDiv1.querySelectorAll("h1")
let testStrong = testDiv1.querySelector("h1>strong")
assert.lengthOf(testDiv1.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

output.innerText = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
const testDiv2 = document.createElement("div");
testDiv2.innerHTML = output.innerText;
headings = testDiv2.querySelectorAll("h1")
testStrong = testDiv2.querySelector("h1>strong")
assert.lengthOf(testDiv2.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `# **título 1**` o `# __título 1__`, debes establecer el HTML interno de `#preview` en `<h1><strong>título 1</strong></h1>`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "# **title 1**";
input.dispatchEvent(new Event("input"));
let headings = preview.querySelectorAll("h1")
let testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");

preview.innerHTML = "";
input.value = "# __title 1__";
input.dispatchEvent(new Event("input"));
headings = preview.querySelectorAll("h1")
testStrong = preview.querySelector("h1>strong")
assert.lengthOf(preview.children, 1);
assert.lengthOf(headings, 1);
assert.lengthOf(headings[0].children, 1);
assert.exists(testStrong);
assert.equal(testStrong.innerText, "title 1");
```

Cuando el valor de `#markdown-input` es `![texto-alt](fuente-imagen)`, `convertMarkdown()` debe devolver `<img alt="texto-alt" src="fuente-imagen">`.

```js
const input = document.querySelector("#markdown-input");
input.value = "![alt-text](image-source)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Cuando el valor de `#markdown-input` es `![texto-alt](fuente-imagen)`, `<img alt="texto-alt" src="fuente-imagen">` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Cuando el valor de `#markdown-input` es `![texto-alt](fuente-imagen)`, `<img alt="texto-alt" src="fuente-imagen">` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "![alt-text](image-source)";
input.dispatchEvent(new Event("input"));
const imgs = preview.querySelectorAll("img");
assert.lengthOf(preview.children, 1)
assert.lengthOf(imgs, 1);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
```

Cuando el valor de `#markdown-input` es `![texto-alt](fuente-imagen)` seguido por `![texto-alt-2](fuente-imagen-2)` en una nueva línea, `convertMarkdown()` debe devolver `<img alt="texto-alt" src="fuente-imagen"><img alt="texto-alt-2" src="fuente-imagen-2">`.

```js
const input = document.querySelector("#markdown-input");
input.value = "![alt-text](image-source)\n![alt-text-2](image-source-2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const imgs = testDiv.querySelectorAll("img");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(imgs, 2);
assert.equal(imgs[0].alt, "alt-text");
assert.isTrue(imgs[0].src.endsWith("image-source"));
assert.equal(imgs[1].alt, "alt-text-2");
assert.isTrue(imgs[1].src.endsWith("image-source-2"));
```

Cuando el valor de `#markdown-input` es `[texto del enlace](URL)`, `convertMarkdown()` debe devolver `<a href="URL">texto del enlace</a>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Cuando el valor de `#markdown-input` es `[texto del enlace](URL)`, `<a href="URL">texto del enlace</a>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Cuando el valor de `#markdown-input` es `[texto del enlace](URL)`, `<a href="URL">texto del enlace</a>` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "[link text](URL)";
input.dispatchEvent(new Event("input"));
const anchors = preview.querySelectorAll("a");
assert.lengthOf(preview.children, 1);
assert.lengthOf(anchors, 1);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
```

Cuando el valor de `#markdown-input` es `[texto del enlace](URL)` seguido por `[texto del enlace 2](URL2)` en una nueva línea, `convertMarkdown()` debe devolver `<a href="URL">texto del enlace</a><a href="URL2">texto del enlace 2</a>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "[link text](URL)\n[link text 2](URL2)";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const anchors = testDiv.querySelectorAll("a");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(anchors, 2);
assert.isTrue(anchors[0].href.endsWith("URL"));
assert.equal(anchors[0].innerText, "link text");
assert.isTrue(anchors[1].href.endsWith("URL2"));
assert.equal(anchors[1].innerText, "link text 2");
```

Cuando el valor de `#markdown-input` es `> esto es una cita`, `convertMarkdown()` debe devolver `<blockquote>esto es una cita</blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Cuando el valor de `#markdown-input` es `> esto es una cita`, `<blockquote>esto es una cita</blockquote>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Cuando el valor de `#markdown-input` es `> esto es una cita`, `<blockquote>esto es una cita</blockquote>` debe añadirse como hijo de `#preview`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> this is a quote";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.equal(quotes[0].innerText, "this is a quote");
```

Cuando el valor de `#markdown-input` es `> esto es una cita` seguido por `> esto es otra cita` en una nueva línea, `convertMarkdown()` debe devolver `<blockquote>esto es una cita</blockquote><blockquote>esto es otra cita</blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> this is a quote\n> this is another quote";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
assert.lengthOf(testDiv.children, 2);
assert.lengthOf(quotes, 2);
assert.equal(quotes[0].innerText, "this is a quote");
assert.equal(quotes[1].innerText, "this is another quote");
```


Cuando el valor de `#markdown-input` es `algún texto > ya no es una cita`, `convertMarkdown()` no debe convertir `> ya no es una cita` en un elemento `blockquote`.

```js
const input = document.querySelector("#markdown-input");
input.value = "some text > not a quote anymore";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const testQuote = testDiv.querySelector("blockquote")
assert.notExists(testQuote);
```

Cuando el valor de `#markdown-input` es `> **esto es una *cita***`, `convertMarkdown()` debe devolver `<blockquote><strong>esto es una <em>cita&lt/em>&lt/strong>&lt/blockquote>`.

```js
const input = document.querySelector("#markdown-input");
input.value = "> **this is a *quote***";
const testDiv = document.createElement("div");
testDiv.innerHTML = convertMarkdown();
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Cuando el valor de `#markdown-input` es `> **esto es una *cita***`, `<blockquote><strong>esto es una <em>cita&lt/em>&lt/strong>&lt/blockquote>` debe mostrarse dentro de `#html-output`.

```js
const input = document.querySelector("#markdown-input");
const output = document.querySelector("#html-output");
output.innerText = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const testDiv = document.createElement("div");
testDiv.innerHTML = output.innerText;
const quotes = testDiv.querySelectorAll("blockquote");
const strongs = testDiv.querySelectorAll("blockquote>strong");
const italics = testDiv.querySelectorAll("blockquote>strong>em");
assert.lengthOf(testDiv.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Cuando el valor de `#markdown-input` es `> **esto es una *cita***`, debes establecer el HTML interno de `#preview` en `<blockquote><strong>esto es una <em>cita&lt/em>&lt/strong>&lt/blockquote>`.

```js
const input = document.querySelector("#markdown-input");
const preview = document.querySelector("#preview");
preview.innerHTML = "";
input.value = "> **this is a *quote***";
input.dispatchEvent(new Event("input"));
const quotes = preview.querySelectorAll("blockquote");
const strongs = preview.querySelectorAll("blockquote>strong");
const italics = preview.querySelectorAll("blockquote>strong>em");
assert.lengthOf(preview.children, 1);
assert.lengthOf(quotes, 1);
assert.lengthOf(quotes[0].children, 1);
assert.lengthOf(strongs, 1);
assert.lengthOf(strongs[0].children, 1);
assert.equal(strongs[0].innerText, "this is a quote");
assert.lengthOf(italics, 1);
assert.equal(italics[0].innerText, "quote");
```

Debes tener sólo un elemento `script` en tu HTML.

```js
assert.lengthOf(document.querySelectorAll("script"), 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <div id="container">
        <div class="container">
            <h2>Markdown Input:</h2>
            <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
        </div>
        <div class="container">
            <h2>Raw HTML Output:</h2>
            <div id="html-output"></div>
        </div>
        <div class="container">
            <h2>HTML Preview:</h2>
            <div id="preview"></div>
        </div>
    </div>
</body>

</html>
```

```css
* {
     box-sizing: border-box;
}
 body {
     font-family: Arial, sans-serif;
     padding: 20px;
}
 #markdown-input {
     width: 100%;
     height: 100px;
}
 #html-output, #preview {
     height: 100px;
     display: inline-block;
     width: 100%;
     border: 1px solid #ccc;
     padding: 10px;
     margin: auto;
     white-space: pre-wrap;
     background-color: #f9f9f9;
}
 @media (min-width: 600px) {
     #markdown-input, #html-output, #preview {
         height: 200px;
         margin: 0;
    }
     #container {
         display: flex;
         justify-content: space-evenly;
         gap: 10px;
    }
}
```

```js

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Markdown to HTML Converter</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <h1>Markdown to HTML Converter</h1>
    <div id="container">
        <div class="container">
            <h2>Markdown Input:</h2>
            <textarea id="markdown-input" placeholder="Enter your markdown here..."></textarea>
        </div>
        <div class="container">
            <h2>Raw HTML Output:</h2>
            <div id="html-output"></div>
        </div>
        <div class="container">
            <h2>HTML Preview:</h2>
            <div id="preview"></div>
        </div>
    </div>
    <script src="script.js"></script>
</body>

</html>
```

```css
* {
     box-sizing: border-box;
}
 body {
     font-family: Arial, sans-serif;
     padding: 20px;
}
 #markdown-input {
     width: 100%;
     height: 100px;
}
 #html-output, #preview {
     height: 100px;
     display: inline-block;
     width: 100%;
     border: 1px solid #ccc;
     padding: 10px;
     margin: auto;
     white-space: pre-wrap;
     background-color: #f9f9f9;
}
 @media (min-width: 600px) {
     #markdown-input, #html-output, #preview {
         height: 200px;
         margin: 0;
    }
     #container {
         display: flex;
         justify-content: space-evenly;
         gap: 10px;
    }
}
```

```js
const patternReplacementArray = [
    { h6: [/^[ \t]*######[ \t](.+$)/gm, "<h6>$1</h6>"] },
    { h5: [/^[ \t]*#####[ \t](.+$)/gm, "<h5>$1</h5>"] },
    { h4: [/^[ \t]*####[ \t](.+$)/gm, "<h4>$1</h4>"] },
    { h3: [/^[ \t]*###[ \t](.+$)/gm, "<h3>$1</h3>"] },
    { h2: [/^[ \t]*##[ \t](.+$)/gm, "<h2>$1</h2>"] },
    { h1: [/^[ \t]*#[ \t](.+$)/gm, "<h1>$1</h1>"] },
    { hr: [/(?<=\n|^)---(?=\n|$)/, "<hr>"] },
    { strong: [/((?:\*|_){2})(.*?)\1/g, "<strong>$2</strong>"] },
    { em: [/(\*|_)(.*?)\1/g, "<em>$2</em>"] },
    { img: [/!\[\s*(.*?)\s*\]\(\s*(.*?)\s*\)/g, '<img alt="$1" src="$2">'] },
    { a: [/\[(.*?)\]\(\s*(.*?)\s*\)/g, '<a href="$2">$1</a>'] },
    { blockquote1: [/^[ \t]*> ([^\n]+?)$/gm, "<blockquote>$1</blockquote>\n"] },
    { blockquote2: [/(?<=<blockquote>)(.+?)<\/blockquote>\n\n^([^<\n]+?)$/gm, "$1 $2</blockquote>\n"] },
    { ul: [/(?:^[-*+]\s+.*?$\n)+\s*?/gm, "<ul>\n$&</ul>"] },
    { ol: [/(?:^\d*\.\s+?\w.*?$\n)+\s*?/gm, "<ol>\n$&</ol>"] },
    { uli: [/^[-*+]\s+(.*?$)\n/gm, "<li>$1</li>\n"] },
    { oli: [/^\d*\.\s+?(.*?$)\n/gm, "<li>$1</li>\n"] },
    { p1: [/^([^<>\s].*?)(?=\n\s*$)/gsm, "<p>$1</p>"] },
    { p2: [/^\w*?(?:<strong>|<em>).+(?:<\/strong>|<\/em>)\s*\w*?[ \t]*$/gm, "<p>$&</p>"] }
]

const markdownInput = document.getElementById('markdown-input');
const convertMarkdown = () => {
    let markdown = markdownInput.value;
    let matchUList = false;
    let matchOList = false;
    patternReplacementArray.forEach((item) => {
        const [key] = Object.keys(item);
        const [pattern, replacement] = Object.values(item)[0];
        if (key === "blockquote2") {
            while (markdown.match(pattern)) {
                markdown = markdown.replace(pattern, replacement);
            }
        } else if (key === "ul") {
            if (markdown.match(pattern)) {
                matchUList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "ol") {
            if (markdown.match(pattern)) {
                matchOList = true;
                markdown = markdown.replace(pattern, replacement);
            }

        } else if (key === "uli" && !matchUList) {
            return;
        } else if (key === "oli" && !matchOList) {
            return;
        } else {
            markdown = markdown.replace(pattern, replacement);
        }
    })
    return markdown;
}

markdownInput.addEventListener('input', () => {
    const html = convertMarkdown();
    const out = document.getElementById('html-output');
    const preview = document.getElementById('preview');
    out.textContent = html;
    preview.innerHTML = html;
});
```
