---
id: 614385513d91ae5c251c2052
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

أضف عنصر `title` بالنص `Magazine`، وعنصر `link` لورقة أنماط الخطوط من `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`، وعنصر `link` لورقة أنماط FontAwesome من `https://use.fontawesome.com/releases/v5.8.2/css/all.css`، وعنصر `link` لورقة الأنماط الخاصة بك `./styles.css`.

ستحمّل ورقة أنماط الخطوط ثلاثة خطوط منفصلة: `Anton` و`Baskervville` و`Raleway`.

# --hints--

يجب أن يحتوي شفرتك على ثلاثة عناصر `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 3);
```

يجب أن تكون عناصر `link` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link')?.length, 3);
```

يجب أن تحتوي عناصر `link` الثلاثة على خاصية `rel` بقيمة `stylesheet`.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

يجب أن يحتوي أحد عناصر الرابط على خاصية `href` مضبوطة على `https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap'));
```

يجب أن يحتوي أحد عناصر الرابط على خاصية `href` مضبوطة على `https://use.fontawesome.com/releases/v5.8.2/css/all.css`.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link.getAttribute('href') === 'https://use.fontawesome.com/releases/v5.8.2/css/all.css'));
```

يجب أن يحتوي أحد عناصر `link` على خاصية `href` مضبوطة على `styles.css`.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
assert.strictEqual(styleElement?.getAttribute('data-href'), 'styles.css');
```

يجب أن يحتوي شفرتك على عنصر `title`.

```js
const title = document.querySelector('title');
assert.isNotNull(title);
```

يجب أن يحتوي شفرتك على عنصر `title` واحد فقط.

```js
const title = document.querySelectorAll('title');
assert.strictEqual(title?.length, 1);
```

يجب أن يكون عنصر `title` داخل عنصر `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('title'));
```

يجب أن يكون لمشروعك عنوان `Magazine`. تذكر أن حالة الأحرف والإملاء مهمان للعنوان.

```js
const title = document.querySelector('title');
assert.equal(title?.text?.trim(), 'Magazine');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```

```css

```
