---
id: 6143869bb45bd85e3b1926aa
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

داخل عنصر `.heading`، أنشئ عنصر `header` مع تعيين `class` إلى `hero`.

داخل ذلك العنصر، أنشئ عنصر `img` مع تعيين `src` إلى `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`، و`alt` إلى `freecodecamp logo`، و`class` إلى `hero-img`.

يمكن تعيين الخاصية `loading` على عنصر `img` إلى `lazy` لإخبار المتصفح بعدم جلب مورد الصورة حتى يحتاجه (أي عندما يقوم المستخدم بالتمرير لرؤية الصورة). كميزة إضافية، لن تُحمّل العناصر التي تستخدم التحميل الكسول حتى تُحمّل العناصر غير الكسولة - وهذا يعني أن المستخدمين ذوي الاتصالات البطيئة بالإنترنت يمكنهم عرض محتوى الصفحة دون انتظار تحميل الصور.

أضف إلى عنصر `img` الجديد الخاص بك خاصية `loading` مع تعيينها إلى `lazy`.

# --hints--

يجب أن تنشئ عنصر `header`.

```js
assert.exists(document.querySelector('header'));
```

يجب أن يكون عنصر `header` داخل عنصر `.heading` الخاص بك.

```js
assert(document.querySelector('header')?.parentElement?.className === 'heading');
```

يجب أن يحتوي عنصر `header` على `class` مع تعيينه إلى `hero`.

```js
assert(document.querySelector('header')?.className === 'hero');
```

يجب أن تنشئ عنصر `img`.

```js
assert.exists(document.querySelector('img'));
```

يجب أن يكون عنصر `img` داخل عنصر `header` الخاص بك.

```js
assert(document.querySelector('img')?.parentElement?.localName === 'header');
```

يجب أن يحتوي عنصر `img` على `src` مع تعيينه إلى `https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png`.

```js
assert(document.querySelector('img')?.getAttribute('src') === 'https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png');
```

يجب أن يحتوي عنصر `img` على `alt` مع تعيينه إلى `freecodecamp logo`.

```js
assert(document.querySelector('img')?.getAttribute('alt') === 'freecodecamp logo');
```

يجب أن يحتوي عنصر `img` على الخاصية `loading` مع تعيينها إلى `lazy`.

```js
assert(document.querySelector('img')?.getAttribute('loading') === 'lazy');
```

يجب أن يحتوي عنصر `img` على `class` مع تعيينه إلى `hero-img`.

```js
assert(document.querySelector('img')?.className === 'hero-img');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
