---
id: 61437d575fb98f57fa8f7f36
title: الخطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

ابدأ بقالب HTML الجاهز القياسي. أضف إعلان `DOCTYPE`، وعنصر `html` يحدد أن الصفحة باللغة الإنجليزية، وعنصر `head`، وعنصر `body`.

أضف وسم `<meta>` مع `charset` المناسب ووسم `<meta>` للاستجابة على الأجهزة المحمولة داخل عنصر `head`.

# --hints--

يجب أن تحتوي شفرتك على مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

يجب أن تضع مسافة بعد مرجع `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

يجب أن تحدد نوع المستند ليكون `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

يجب أن تغلق إعلان `DOCTYPE` بعلامة `>` بعد النوع.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

يجب أن يكون لديك وسم افتتاحي `<html>`، ويجب أن يحتوي على `lang` بقيمة `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert(code.match(/<\/html\s*>/));
```

يجب أن يكون إعلان `DOCTYPE` في بداية شفرة HTML الخاصة بك.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

يجب أن يكون لديك وسم افتتاحي `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

يجب أن يكون لديك وسم افتتاحي `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

يجب أن يكون لديك وسم إغلاق `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

يجب أن يكون عنصرا `head` و `body` أخوين (متجاورين).

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

يجب أن يكون عنصر `head` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

يجب أن يكون عنصر `body` داخل عنصر `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

يجب أن يكون لديك وسمان `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

يجب أن يحتوي وسم `meta` واحد على خاصية `name` بقيمة `viewport`، و`content` بقيمة `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

يجب أن يحتوي وسم `meta` الآخر على خاصية `charset` بقيمة `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
