---
id: 614387cbefeeba5f3654a291
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

بعد عنصر `header`، أنشئ عنصر `div` مع تعيين `class` إلى `author`.

داخل ذلك العنصر `div`، أنشئ عنصر `p` مع تعيين `class` إلى `author-name` وأضف له النص `By freeCodeCamp`. غلف جزء `freeCodeCamp` بعنصر `a` مع تعيين `href` إلى `https://freecodecamp.org`، و`target` إلى `_blank`.


أسفل ذلك، أضف عنصر `p` ثانيًا مع تعيين الصنف `publish-date` والنص `March 7, 2019`.

# --hints--

يجب أن تنشئ عنصر `div` جديد.

```js
assert.exists(document.querySelector('div'));
```

يجب أن يأتي عنصر `div` بعد عنصر `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

يجب أن يحتوي عنصر `div` على `class` يساوي `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

يجب أن تنشئ عنصرين جديدين من نوع `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

يجب أن يكون العنصران الجديدان `p` داخل عنصر `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

يجب أن يحتوي أول عنصر `p` جديد على `class` يساوي `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

يجب أن يحتوي أول عنصر `p` جديد على النص `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

يجب أن يحتوي ثاني عنصر `p` جديد على `class` يساوي `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

يجب أن يحتوي ثاني عنصر `p` جديد على النص `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

يجب أن تنشئ عنصر `a` جديد.

```js
assert.exists(document.querySelector('a'));
```

يجب أن يكون عنصر `a` داخل أول عنصر `p` جديد.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

يجب أن يحتوي عنصر `a` على `href` يساوي `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

يجب أن يحتوي عنصر `a` على `target` يساوي `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

يجب أن يغلف عنصر `a` النص `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
