---
id: 6143862a5e54455d262c212e
title: خطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

داخل `body`، أنشئ عنصر `main`. ثم داخل هذا العنصر، أنشئ `section` مع تعيين `class` إلى `heading`.

# --hints--

يجب أن يكون لديك عنصر `main`.

```js
assert.exists(document.querySelector('main'));
```

يجب أن يكون عنصر `main` داخل عنصر `body`.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

يجب أن يكون لديك عنصر `section`.

```js
assert.exists(document.querySelector('section'));
```

يجب أن يكون عنصر `section` داخل عنصر `main`.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

يجب أن يحتوي عنصر `section` على `class` معين إلى `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
