---
id: 68e4c81577243e4db1676491
title: خطوة 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

في هذه الورشة، ستتدرّب على العمل مع HTML الدلالي من خلال بناء صفحة ويب تتضمن بعض نصائح Quincy Larson للحصول على وظيفة مطوِّر. تم إعداد قالب جاهز أساسي لـ HTML من أجلك.

ابدأ بإنشاء عنصر `h1` يحتوي على النص `Quincy's Tips for Getting a Developer Job`.

# --hints--

يجب أن يكون لديك عنصر `h1`.

```js
assert.exists(document.querySelector('h1'));
```

يجب أن يحتوي عنصر `h1` الخاص بك على النص `Quincy's Tips for Getting a Developer Job`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), "Quincy's Tips for Getting a Developer Job")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
