---
id: 68e65ec5348e40a739903e6d
title: الخطوة 2
challengeType: 0
dashedName: step-2
---

# --description--

الآن، أنشئ عنصر فقرة أسفل عنصر `h1`. داخل هذه الفقرة أضف النص `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

# --hints--

يجب أن يكون لديك عنصر `p`.

```js
assert.exists(document.querySelector('p'));
```

يجب أن تحتوي الفقرة على النص `Learning to code is hard, but as Quincy Larson says, You can become a developer.`. تحقق من صحة الإملاء.

```js
assert.equal(document.querySelector('p')?.textContent.trim(), "Learning to code is hard, but as Quincy Larson says, You can become a developer.");
```

يجب أن تكون الفقرة أسفل عنصر `h1`.

```js
assert.exists(document.querySelector('h1 + p'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
