---
id: 68e90335bd53a901c41dfc13
title: الخطوة 3
challengeType: 0
dashedName: step-3
---

# --description--

في درس سابق تعلمت أن العنصر `q` يُستخدم لتمييز النص المقتبس عن المحتوى المحيط به.

إليك مثالًا:

```html
<p>
  Nancy said, <q>Learning is fun!</q>
</p>
```

معظم المتصفحات الحديثة تضيف علامات اقتباس تلقائيًا حول الاقتباس المضمَّن عند استخدامك العنصر `q`.

داخل الفقرة، غلف النص `You can become a developer.` بعنصر اقتباس مضمَّن، مع إبقاء بقية الفقرة دون تغيير.

# --hints--

يجب أن يكون لديك عنصر `q`.

```js
assert.exists(document.querySelector('q'));
```

يجب أن يكون عنصر `q` مسكنًا داخل الفقرة التي تقع تحت العنصر `h1`.

```js
assert.exists(document.querySelector('h1 + p > q'));
```

يجب أن يحتوي عنصر `q` على النص `You can become a developer.`

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent.trim(), 'You can become a developer.');
```

داخل عنصر `q`، يجب ألا يحتوي النص على فراغات بادئة أو لاحقة (لا مسافات أو علامات تبويب أو فواصل أسطر قبل الحرف الأول أو بعد الأخير).

```js
assert.equal(document.querySelector('h1 + p > q')?.textContent, 'You can become a developer.');
```

يجب إضافة وسم `q` فقط حول الجزء المناسب من الجملة دون تغيير النص الأصلي. يجب أن تظل فقرة `p` تقرأ: `Learning to code is hard, but as Quincy Larson says, You can become a developer.`

```js
const pElCleanedText = document.querySelector('h1 + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, 'Learning to code is hard, but as Quincy Larson says, You can become a developer.');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
--fcc-editable-region--
    <p>Learning to code is hard, but as Quincy Larson says, You can become a developer.</p>
--fcc-editable-region--
  </body>
</html>
```
