---
id: 68e96ce52087043b0e999296
title: الخطوة 5
challengeType: 0
dashedName: step-5
---

# --description--

أسفل عنصر الفقرة، أضف عنصر `main` وقم بالتسكين لثلاثة عناصر `section` بداخله.

# --hints--

يجب أن يكون لديك عنصر `main` أسفل عنصر الفقرة.

```js
const mainEl = document.querySelector('p + main');
assert.exists(mainEl);
```

يجب أن يكون لديك ثلاثة عناصر `section` داخل عنصر `main`.

```js
const sections = document.querySelectorAll('main > section');
assert.lengthOf(sections, 3);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
