---
id: 68e964b84f954c322376e592
title: الخطوة 4
challengeType: 0
dashedName: step-4
---

# --description--

إذا كان مصدر الاقتباس موقع ويب، يمكنك الاستشهاد به باستخدام خاصية `cite`. يجب أن تكون قيمة هذه الخاصية عنوان URL صالحًا. رغم أن هذه الخاصية لا تغير طريقة عرض الاقتباس المضمّن، إلا أنها مفيدة جدًا لمنح قارئي الشاشة ومحركات البحث معلومات إضافية عن الاقتباس.

إليك مثال على عنصر اقتباس مضمّن مع خاصية `cite`:

```html
<p>
  Nancy said,
  <q cite="https://example.com">Learning is fun!</q>
</p>
```

أضف خاصية `cite` إلى عنصر الاقتباس المضمّن مع هذا العنوان:

`https://www.freecodecamp.org/news/learn-to-code-book/`

# --hints--

يجب أن يحتوي عنصر `q` الخاص بك على خاصية `cite`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.exists(qEl?.getAttribute('cite'));
```

يجب أن تكون قيمة خاصية `cite` في عنصر `q` الخاص بك هي `https://www.freecodecamp.org/news/learn-to-code-book/`.

```js
const qEl = document.querySelector('h1 + p > q');
assert.equal(qEl?.getAttribute('cite'), 'https://www.freecodecamp.org/news/learn-to-code-book/');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
--fcc-editable-region--
      <q>You can become a developer.</q>
--fcc-editable-region--
    </p>
  </body>
</html>
```
