---
id: 68eaa9ef8a0867071f195f59
title: الخطوة 10
challengeType: 0
dashedName: step-10
---

# --description--

حتى الآن كنت تستخدم الخاصية `cite` لنسب مصدر الاقتباس، لكن الخاصية لا تعرض المصدر للمستخدم فعليًا.

إذا أردت نسب المصدر بصريًا، يمكنك إضافة عنصر اقتباس، `cite`، خارج عنصر الاقتباس المقطعي. عنصر الاقتباس هو عنصر HTML يمكنك استخدامه لوضع عنوان عمل إبداعي مرجعي، مثل كتاب أو مقال أو أغنية أو فيلم أو موقع ويب أو ورقة بحثية.

إليك مثالًا:

```html
<div>
  <blockquote cite="https://www.freecodecamp.org/news/is-college-worth-it/">
    The first thing you should consider about education is this is an economic decision.
  </blockquote>
  <p>&mdash;Quincy Larson, <cite>Is College Still Worth it? Tips from my 20 Years in Adult Education</cite></p>
</div>
```

داخل عنصر `p` أسفل عنصر الاقتباس المقطعي، غلف `How to Learn to Code and Get a Developer Job [Full Book]` بعنصر `cite`.

# --hints--

يجب أن يكون لديك عنصر `cite` داخل عنصر `p` أسفل عنصر `blockquote`.

```js
assert.exists(document.querySelector('blockquote + p > cite'));
```

يجب أن يحتوي عنصر `cite` على النص `How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const citeEl = document.querySelector('blockquote + p > cite');
assert.equal(citeEl?.innerText.trim(), 'How to Learn to Code and Get a Developer Job [Full Book]');
```

باستثناء عنصر `cite` الذي أضفته للتو، يجب أن يظل نص عنصر `p` هو `—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]`.

```js
const pElCleanedText = document.querySelector('blockquote + p')?.innerText.trim().replaceAll(/ {2,}/g, ' ');
assert.equal(pElCleanedText, '—Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        <blockquote cite="https://www.freecodecamp.org/news/learn-to-code-book/">
          Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?
        </blockquote>
        <p>
      --fcc-editable-region--
          &mdash;Quincy Larson, How to Learn to Code and Get a Developer Job [Full Book]
      --fcc-editable-region--
        </p>
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
