---
id: 68ea7e7b916e8270483689b6
title: الخطوة 7
challengeType: 0
dashedName: step-7
---

# --description--

أسفل عنصر `h2`، أضف اقتباسًا آخر من كوينسي. هذه المرة، لن يكون الاقتباس جزءًا من فقرة أكبر. بدلًا من ذلك، ستكون الفقرة بأكملها اقتباسًا. لتمييز النص المقتبس بهذا الشكل، يجب استخدام عنصر الاقتباس المضمّن: `blockquote`. في المتصفح، سترى أن النص مُزاح قليلًا إلى الداخل.

إليك مثالًا على عنصر الاقتباس المضمّن مع نص مقتبس:

```html
<blockquote>
  The first thing you should consider about education is this is an economic decision.
</blockquote>
```

الآن، داخل القسم الأول، أضف عنصر اقتباس مضمّن أسفل عنصر `h2` بالنص `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

# --hints--

داخل عنصر `section` الأول، يجب أن يكون لديك عنصر `blockquote` أسفل عنصر `h2`.

```js
assert.exists(document.querySelector('section:first-of-type > h2 + blockquote'));
```

يجب أن يحتوي عنصر `blockquote` الخاص بك على النص `Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?`.

```js
const firstBlockquoteEl = document.querySelector('section:first-of-type > h2 + blockquote');
assert.equal(firstBlockquoteEl?.innerText.trim(), 'Can you imagine what it would be like to be a successful developer? To have built software systems that people rely upon?');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Quincy's Tips for Getting a Developer Job</title>
  </head>
  <body>
    <h1>Quincy's Tips for Getting a Developer Job</h1>
    <p>
      Learning to code is hard, but as Quincy Larson says,
      <q cite="https://www.freecodecamp.org/news/learn-to-code-book/">You can become a developer.</q>
    </p>

    <main>
      <section>
        <h2>Envisioning Success</h2>
        --fcc-editable-region--
        
        --fcc-editable-region--
      </section>
      <section>

      </section>
      <section>

      </section>
    </main>
  </body>
</html>
```
