---
id: af7588ade1100bde429baf20
title: Costruisci un rilevatore di lettere mancanti
challengeType: 26
dashedName: build-a-missing-letter-detector
---

# --description--

**Obiettivo:** Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione chiamata `fearNotLetter`.
1. La funzione `fearNotLetter` dovrebbe accettare un argomento: una stringa che rappresenta un intervallo di lettere in ordine alfabetico che può avere una lettera mancante.
1. La funzione dovrebbe trovare la lettera mancante nell'intervallo di lettere passato e restituirla.
1. Se tutte le lettere sono presenti nell'intervallo, la funzione dovrebbe restituire `undefined`.

# --hints--

Dovresti avere una funzione `fearNotLetter`.

```js
assert.isFunction(fearNotLetter);
```

`fearNotLetter("abce")` dovrebbe restituire la stringa `d`.

```js
assert.deepEqual(fearNotLetter('abce'), 'd');
```

`fearNotLetter("abcdefghjklmno")` dovrebbe restituire la stringa `i`.

```js
assert.deepEqual(fearNotLetter('abcdefghjklmno'), 'i');
```

`fearNotLetter("stvwx")` dovrebbe restituire la stringa `u`.

```js
assert.deepEqual(fearNotLetter('stvwx'), 'u');
```

`fearNotLetter("bcdf")` dovrebbe restituire la stringa `e`.

```js
assert.deepEqual(fearNotLetter('bcdf'), 'e');
```

`fearNotLetter("abcdefghijklmnopqrstuvwxyz")` dovrebbe restituire `undefined`.

```js
assert.isUndefined(fearNotLetter('abcdefghijklmnopqrstuvwxyz'));
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function fearNotLetter (str) {
  for (var i = str.charCodeAt(0); i <= str.charCodeAt(str.length - 1); i++) {
    let letter = String.fromCharCode(i);
    if (str.indexOf(letter) === -1) {
      return letter;
    }
  }

  return undefined;
}
```
