---
id: 65c4f2851645e8d84d14f5d1
title: Passo 11
challengeType: 20
dashedName: step-11
---

# --description--

Agora você precisa percorrer a árvore recursivamente e inserir os valores usando o princípio para árvores binárias:

- Valores menores que a chave são colocados na subárvore esquerda
- Valores maiores que a chave são colocados na subárvore direita

Após sua declaração condicional existente, escreva outra declaração `if` para verificar se `key` é menor que `node.key`.

# --hints--

Você deve escrever outra instrução if para verificar se `key` é menor que `node.key`. Lembre-se de usar a palavra-chave `pass`.

```js
({ test: () => assert(runPython(` _Node(_code).find_class("BinarySearchTree").find_function("_insert").find_ifs()[1].find_conditions()[0].is_equivalent("key < node.key")`)) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

--fcc-editable-region--
```

