---
id: 65c4f2b7178afed8e88f782f
title: Passo 12
challengeType: 20
dashedName: step-12
---

# --description--

Se `key` < `node.key` for avaliado como `True`, então o novo nó deve ser colocado na subárvore esquerda.

Delete `pass` e chame recursivamente o método `_insert` com o filho esquerdo como o primeiro argumento e `key` como o segundo argumento. Atribua o resultado ao atributo `left` do nó atual.

# --hints--

Você deve remover a palavra-chave `pass` do bloco `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Você deve chamar o método `self._insert` passando `node.left` e `key` como argumentos.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Você deve atribuir o resultado da sua chamada `self._insert()` ao atributo `left` do nó atual.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
