---
id: 65c4f09e074dd8d37830ea00
title: Passo 6
challengeType: 20
dashedName: step-6
---

# --description--

Crie outra classe vazia chamada `BinarySearchTree` que representa uma árvore binária de busca.

# --hints--

Você deve ter uma classe chamada `BinarySearchTree`. Lembre-se de usar a palavra-chave `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
