---
id: 65c4f0b82db314d3fcc0b8e1
title: Passo 7
challengeType: 20
dashedName: step-7
---

# --description--

Dentro da classe `BinarySearchTree`, substitua `pass` por um método `__init__` e adicione um parâmetro `self` a este método.

# --hints--

Você deve remover a palavra-chave `pass` da classe `BinarySearchTree`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


Você deve definir um método `__init__` e adicionar um parâmetro `self` a este método.Lembre-se de usar a palavra-chave `pass` dentro do corpo do método.

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
        
--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
