---
id: 65c4f22498d22ed775ef8efb
title: Passo 10
challengeType: 20
dashedName: step-10
---

# --description--

Agora você precisa verificar se o parâmetro `node` é `None`.Se for, isso significa que o método alcançou um nó folha ou um local vazio na árvore onde o novo nó deve ser inserido.

Dentro do corpo do método `_insert`, substitua `pass` por uma declaração `if` que verifica se `node is None`.

Dentro do novo bloco `if`, retorne `TreeNode(key)` para criar uma nova instância de `TreeNode` com a chave fornecida.Este se tornará o novo nó folha, inserindo efetivamente a chave na árvore.

# --hints--

Você não deve ter `pass` no seu método `_insert`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Você deve escrever uma declaração `if` para verificar se `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Você deve retornar `TreeNode(key)` do bloco `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
