---
id: 65c4f013851cefd1a4fe4c96
title: Passo 3
challengeType: 20
dashedName: step-3
---

# --description--

O método `__init__` recebe dois parâmetros: `self` (que representa a instância da classe sendo criada) e `key` (o valor a ser armazenado no nó). Adicione esses dois parâmetros ao método `__init__()`.

# --hints--

Você deve adicionar o parâmetro `self` ao método.

```js
({ test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*/) });
```

Você deve adicionar o parâmetro `key` ao método.

```js
({
  test: () => assert.match(code, /def\s+__init__\s*\(\s*self\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__():
        pass
--fcc-editable-region--
```
