---
id: 65c4f0d32ca17ad4b1636b0e
title: Passo 8
challengeType: 20
dashedName: step-8
---

# --description--

Dentro do método `__init__`, delete `pass` e inicialize `root` com o valor `None`. 

O atributo `root` representa o nó raiz da árvore binária de busca. Como este é o construtor, quando um novo objeto `BinarySearchTree` é criado, ele começa com uma árvore vazia, então o atributo `root` é definido como `None`.

# --hints--

Você deve remover a declaração `pass` do método `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Você deve inicializar o atributo `root` com `None` usando `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
