---
id: 65c4eb814cc977c95cd7df0e
title: Passo 1
challengeType: 20
dashedName: step-1
---

# --description--

Neste projeto, você vai criar uma Binary Search Tree (BST). Uma BST é uma estrutura de dados na qual cada nó tem no máximo dois filhos, com o filho esquerdo contendo valores menores que o nó pai e o filho direito contendo valores maiores que o nó pai, permitindo operações eficientes de busca e ordenação.

É assim que uma Binary Search Tree se parece:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Comece definindo uma classe `TreeNode` vazia. A classe `TreeNode` representa um nó em uma árvore binária de busca. Use a palavra-chave `pass` para preencher o corpo da classe e evitar um erro.

# --hints--

Você deve usar a palavra-chave `class` para declarar uma classe vazia chamada `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
