---
id: 65c4f37ba6cbcfdc77e24165
title: Passo 17
challengeType: 20
dashedName: step-17
---

# --description--

O método `insert` será chamado pelo usuário.Além do parâmetro `self`, ele também precisará de um parâmetro `key`.Este parâmetro será o valor da chave a ser inserido na árvore binária de busca.

Adicione `key` como o segundo parâmetro na definição da função.

# --hints--

O método `insert` deve conter o parâmetro `self`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,/) });
```

O método `insert` deve conter o parâmetro `key`.

```js
({ test: () => assert.match(code, /def\s+insert\s*\(\s*self\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)
        return node

--fcc-editable-region--            
    def insert(self):
        pass

--fcc-editable-region--

```
