---
id: 65c4f300da28d8da361bfa93
title: Passo 14
challengeType: 20
dashedName: step-14
---

# --description--

Dentro da sua cláusula `elif`, substitua `pass` por uma chamada ao método `_insert` com o filho direito do nó atual como o primeiro argumento e `key` como o segundo argumento.Atribua o resultado ao atributo `right` do nó atual.

# --hints--

Você deve remover a palavra-chave `pass` do bloco `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Você deve chamar o método `self._insert()` passando `node.right` e `key` como argumentos.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Você deve atribuir o resultado da sua chamada `self._insert()` ao atributo direito (`node.right`) do nó atual.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
