---
id: 65c4ef8463f869d0eaf87c67
title: Passo 2
challengeType: 20
dashedName: step-2
---

# --description--

Dentro da classe `TreeNode`, substitua `pass` por um método `__init__` para que você possa inicializar os atributos do objeto. Não adicione nenhum parâmetro por enquanto.

# --hints--

Você deve remover a palavra-chave `pass` da classe `TreeNode` e movê-la para dentro do método `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Você deve definir um método `__init__` dentro da sua classe `TreeNode`. Lembre-se de usar a palavra-chave `pass` internamente.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
