---
id: 65c4f02dbd3557d23f12db62
title: Passo 4
challengeType: 20
dashedName: step-4
---

# --description--

Dentro do método `__init__`, delete `pass` e atribua o valor do parâmetro `key` ao atributo `key` do nó usando `self.key`.

Isso significa que o atributo `key` da instância `TreeNode` será definido com o valor passado durante a criação do objeto.

# --hints--

Você deve remover a declaração `pass` do método `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Você deve atribuir o valor do parâmetro `key` ao atributo `key` do nó usando `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
