---
id: 65c4f195de7e2ad5932be717
title: Passo 9
challengeType: 20
dashedName: step-9
---

# --description--

Em seguida, você precisa definir um mecanismo para inserir nós na árvore.Para isso, você precisa definir um método `_insert`, que é uma função auxiliar e será usado pelo método `insert` real mais tarde. 

Este método é recursivo, significando que ele chama a si mesmo para percorrer a árvore até que a localização apropriada para o novo nó seja encontrada.

Defina um método `_insert` com os parâmetros `self`, `node` e `key`.

# --hints--

Você deve definir um método `_insert` dentro da classe `BinarySearchTree`. Lembre-se de usar `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Seu método `_insert` deve receber três parâmetros: `self`, `node` e `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
