---
id: 65c4f07da204f4d2f325fbd3
title: Passo 5
challengeType: 20
dashedName: step-5
---

# --description--

Dentro do método `__init__`, inicialize os atributos `left` e `right` do nó com `None`. Isso ocorre porque quando um nó é criado pela primeira vez, ele não tem filhos à esquerda ou à direita. Lembre-se de usar a palavra-chave `self`. 

# --hints--

Você deve atribuir `None` ao atributo `left` do nó usando `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Você deve atribuir `None` ao atributo `right` do nó usando `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
