---
id: 65c63a4c4da62e9ae18e321a
title: Passo 19
challengeType: 20
dashedName: step-19
---

# --description--

É hora de trabalhar na funcionalidade de busca. Assim como você criou um método auxiliar `_insert` para o método `insert`, você precisa criar um método auxiliar `_search` para o método `search`.

Defina o método `_search` com três parâmetros, a saber, `self`, `node` e `key`.

# --hints--

Você deve declarar um método `_search`. Observe o underscore inicial e lembre-se de usar `pass`.

```js
({ test: () => assert.match(code, /def\s+_search/) });
```

Seu método `_search` deve receber três parâmetros: `self`, `node` e `key`.

```js
({ test: () => assert.match(code, /def\s+_search\(\s*self\s*,\s*node\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
--fcc-editable-region--

--fcc-editable-region--
```
