---
id: 65c646d4148ae3b2d1cbcac4
title: Passo 30
challengeType: 20
dashedName: step-30
---

# --description--

Observe que sua busca retorna algo como `80: <__main__.TreeNode object at 0x108b3e0>`. Esta é a representação padrão em string ao imprimir uma instância de uma classe.

Para alterar isso para imprimir um valor útil, defina outro método chamado `__str__` na classe `TreeNode`. Ele recebe um único argumento `self`.

Após definir `__str__` você receberá uma exceção no console porque o método `__str__` ainda não retorna nada. Você trabalhará no corpo do método na próxima etapa.

# --hints--

Você deve definir um método `__str__` que recebe um único argumento `self`. Lembre-se de usar `pass`.

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
