---
id: 65c6409418806da8d0636ffc
title: Passo 26
challengeType: 20
dashedName: step-26
---

# --description--

As funcionalidades `insert` e `search` estão completas, é hora de testá-las.

Observe que, neste ponto, os nós não estão ordenados e apenas inseridos.Você vai trabalhar na ordenação usando a travessia inorder mais tarde.

Você pode criar uma instância de uma classe em Python assim:

```python
object_name = ClassName()
```

Fora das definições de classe, crie uma instância da classe `BinarySearchTree` e atribua-a à variável `bst`.


# --hints--

Você deve criar uma instância da classe `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Você deve atribuir a nova instância de `BinarySearchTree` à variável `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
