---
id: 65c63fd9b387f0a6c5bf6a72
title: Passo 24
challengeType: 20
dashedName: step-24
---

# --description--

Dentro do método `search`, delete `pass` e chame o método auxiliar `_search` com os seguintes argumentos.

- `self.root`: Esta é a raiz da árvore binária de busca. A busca começa a partir da raiz.
- `key`: Este é o valor que o usuário deseja encontrar na árvore binária de busca.

Internamente, `search` delega a lógica real de busca para o método auxiliar `_search` que realiza a busca recursiva real dentro da árvore binária de busca.

# --hints--

Você deve remover a palavra-chave `pass` do método `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Você deve chamar o método `_search` dentro do método `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Você deve chamar o método `_search` passando `self.root` e `key` como argumentos.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
