---
id: 65c63da2ed6769a10e141341
title: Passo 20
challengeType: 20
dashedName: step-20
---

# --description--

Agora você vai definir um caso base para a busca recursiva.Remova o `pass` atual e escreva uma declaração `if` que verifique duas condições:

- Se `node` for `None`: Isso indica que a busca chegou ao fim de um ramo sem encontrar a chave.
- Se `node.key == key`: Isso significa que a chave foi encontrada no nó atual.

Combine as duas condições com o operador `or` e retorne o nó atual dentro do bloco `if`.

# --hints--

Você deve remover a palavra-chave `pass` do método `_search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_search").has_pass()`
      )
    );
  },
});

```

Você deve escrever uma declaração `if` que verifique se `node` é `None` ou se `node.key` é igual a `key`.

```js
({ test: () => assert.match(code, /^\s{8}if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:/m) })
```

Você deve retornar `node` do bloco `if`.

```js
({ test: () => assert.match(code, /^(\s{8})if\s+node\s+is\s+None\s+or\s+node\.key\s*==\s*key\s*:\s*^\1\s{4}return\s+node/m) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        pass

--fcc-editable-region--
```
