---
id: 65c63e6962159fa372ecdbec
title: Passo 22
challengeType: 20
dashedName: step-22
---

# --description--

Se a segunda declaração `if` não for `True`, isso significa que a chave alvo é maior ou igual à chave do nó atual. 

Em uma árvore de busca binária, se a chave alvo for maior que a chave do nó atual, a busca continua na subárvore direita.

Após o bloco `if`, retorne o resultado da chamada do método `_search` com o filho direito do nó atual e a chave como argumentos.

# --hints--

Você deve chamar recursivamente o método `_search` usando `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Você deve passar dois argumentos para a chamada do método `_search`: `node.right` e `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Você deve retornar o resultado da chamada do método `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
