---
id: 65ca07dd6fa8840491b7a5cd
title: Passo 41
challengeType: 20
dashedName: step-41
---

# --description--

Se nenhuma das condições anteriores for atendida, isso significa que o nó possui filhos à esquerda e à direita.

Para escolher o sucessor, você precisa encontrar o valor mínimo na subárvore direita. O menor valor será o sucessor em ordem do nó atual.

Para encontrar o menor valor, crie uma função auxiliar `_min_value` que recebe dois parâmetros: `self` e `node`.

# --hints--

Você deve definir o método `_min_value` com `self` e `node` como parâmetros. Lembre-se de usar a palavra-chave `pass`.

```js
({ test: () => assert.match(code, /def\s+_min_value\(\s*self\s*,\s*node\s*\)\s*:/) })
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
--fcc-editable-region--
    
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)
    
# print('Search for 80:', bst.search(80))


```
