---
id: 65ca0d5adf39c410cd1177cc
title: Passo 50
challengeType: 20
dashedName: step-50
---

# --description--

Agora, você vai trabalhar na travessia da árvore com base no método de travessia in-order. A travessia in-order é um algoritmo de travessia em profundidade de árvore binária que visita a subárvore esquerda, o nó atual e depois a subárvore direita.

Defina o método `_inorder_traversal` dentro da classe `BinarySearchTree` e dê a ele três parâmetros: `self`, `node` e `result`. Onde `node` é o nó atual sendo considerado durante a travessia e `result` é a lista à qual as chaves são adicionadas em ordem ordenada.

# --hints--

Você deve definir um método `_inorder_traversal` que recebe três parâmetros: `self`, `node` e `result`. Lembre-se de usar `pass`.

```js
assert.match(code, /def\s+_inorder_traversal\(\s*self\s*,\s*node\s*,\s*result\s*\)/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left   
            
            node.key = self._min_value(node.right)
            node.right = self._delete(node.right, node.key)   
        
        return node

    def delete(self, key):
        self.root = self._delete(self.root, key)

    def _min_value(self, node):
        while node.left is not None:
            node = node.left
        return node.key

--fcc-editable-region--

--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
