---
id: 65ca075e3184180309f4c1d9
title: Passo 39
challengeType: 20
dashedName: step-39
---

# --description--

Dentro do bloco `else`, substitua `pass` por uma instrução `if` para verificar se o filho esquerdo do nó atual é `None`.

Quando `node.left` é `None`, não há filho à esquerda. Portanto, retorne o filho direito do novo bloco `if` como substituição.

# --hints--

Você deve remover a palavra-chave `pass` do bloco `else`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree").find_function("_delete"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});
```

Você deve escrever uma declaração `if` que verifique se `node.left` é `None`.

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

Você deve retornar o filho direito do nó atual.

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            pass   
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
