---
id: 587d78af367417b2b2512b03
title: 설문조사 양식 만들기
challengeType: 25
dashedName: build-a-survey-form
demoType: onClick
saveSubmissionToDB: true
---

# --description--

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `h1` 요소에 `id`가 `title`인 페이지 제목이 있어야 합니다.
1. `p` 요소에 `id`가 `description`인 간단한 설명이 있어야 합니다.
1. `form` 요소에 `id`가 `survey-form`인 요소가 있어야 합니다.
1. form 요소 안에 이름을 입력하는 필수 `input` 필드가 있어야 하며, `id`가 `name`이고 `type`가 `text`이어야 합니다.
1. form 요소 안에 이메일을 입력하는 필수 `input` 필드가 있어야 하며, `id`가 `email`이어야 합니다.
1. 올바른 형식이 아닌 이메일을 입력하면 HTML5 유효성 검사 오류가 보여야 합니다.
1. form 요소 안에 숫자를 입력하는 `input` 필드가 있어야 하며, `id`가 `number`이어야 합니다.
1. 숫자 입력 필드는 숫자가 아닌 값을 입력하지 못하게 하거나(브라우저에 따라 다름) HTML5 유효성 검사 오류를 보여야 합니다.
1. 숫자 입력 필드의 `min`와 `max` 속성으로 정의된 범위를 벗어난 숫자를 입력하면 HTML5 유효성 검사 오류가 보여야 합니다.
1. 이름, 이메일, 숫자 입력 필드 각각에 대해 form 안에 해당 필드의 목적을 설명하는 `label` 요소가 있어야 하며, 각각 `name-label`, `email-label`, `number-label` 아이디를 가져야 합니다.
1. 이름, 이메일, 숫자 입력 필드 각각에 대해 설명이나 지침을 제공하는 자리 표시자 텍스트가 있어야 합니다.
1. form 요소 안에 `select` 드롭다운 요소가 있어야 하며, `id`가 `dropdown`이고 최소 두 개 이상의 선택 옵션이 있어야 합니다.
1. form 요소 안에 `name` 속성으로 그룹화된 최소 두 개 이상의 라디오 버튼 그룹에서 옵션을 선택할 수 있어야 합니다.
1. form 요소 안에 `value` 속성이 있는 여러 체크박스에서 여러 필드를 선택할 수 있어야 합니다.
1. form 요소 안에 추가 의견을 위한 `textarea`가 있어야 합니다.
1. form 요소 안에 모든 입력을 제출하는 `id`가 `submit`인 버튼이 있어야 합니다.

# --hints--

`h1` 요소에 `id`가 `title`인 요소가 있어야 합니다.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'H1');
```

`#title`이 비어 있으면 안 됩니다.

```js
const el = document.getElementById('title');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`p` 요소에 `id`가 `description`인 요소가 있어야 합니다.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'P');
```

`#description`이 비어 있으면 안 됩니다.

```js
const el = document.getElementById('description');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`form` 요소에 `id`가 `survey-form`인 요소가 있어야 합니다.

```js
const el = document.getElementById('survey-form');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'FORM');
```

`input` 요소에 `id`가 `name`인 요소가 있어야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#name`에 `type`가 `text`인 속성이 있어야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.strictEqual(el.type, 'text');
```

`#name`은 입력을 필수로 해야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#name`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #name');
assert.isNotNull(el);
```

`input` 요소에 `id`가 `email`인 요소가 있어야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#email`에 `type`가 `email`인 속성이 있어야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.strictEqual(el.type, 'email');
```

`#email`은 입력을 필수로 해야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isTrue(el.required);
```

`#email`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #email');
assert.isNotNull(el);
```

`input` 요소에 `id`가 `number`인 요소가 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'INPUT');
```

`#number`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #number');
assert.isNotNull(el);
```

`#number`에 `type`가 `number`인 속성이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.strictEqual(el.type, 'number');
```

`#number`에 숫자 값이 있는 `min` 속성이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.min);
assert.isTrue(isFinite(el.min));
```

`#number`에 숫자 값이 있는 `max` 속성이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotEmpty(el.max);
assert.isTrue(isFinite(el.max));
```

`label` 요소에 `id`가 `name-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 요소에 `id`가 `email-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`label` 요소에 `id`가 `number-label`인 요소가 있어야 합니다.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'LABEL');
```

`#name-label`에는 입력을 설명하는 텍스트가 포함되어야 합니다.

```js
const el = document.getElementById('name-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#email-label`에는 입력을 설명하는 텍스트가 포함되어야 합니다.

```js
const el = document.getElementById('email-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#number-label`에는 입력을 설명하는 텍스트가 포함되어야 합니다.

```js
const el = document.getElementById('number-label');
assert.isNotNull(el);
assert.isAbove(el.innerText.length, 0);
```

`#name-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #name-label');
assert.isNotNull(el);
```

`#email-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #email-label');
assert.isNotNull(el);
```

`#number-label`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #number-label');
assert.isNotNull(el);
```

`#name`에 `placeholder` 속성과 값이 있어야 합니다.

```js
const el = document.getElementById('name');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#email`에 `placeholder` 속성과 값이 있어야 합니다.

```js
const el = document.getElementById('email');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`#number`에 `placeholder` 속성과 값이 있어야 합니다.

```js
const el = document.getElementById('number');
assert.isNotNull(el);
assert.isNotNull(el.placeholder);
assert.isAbove(el.placeholder.length, 0);
```

`select` 필드에 `id`가 `dropdown`인 속성이 있어야 합니다.

```js
const el = document.getElementById('dropdown');
assert.isNotNull(el);
assert.strictEqual(el.tagName, 'SELECT');
```

`#dropdown`에는 선택 가능(비활성화되지 않은) `option` 요소가 최소 두 개 있어야 합니다.

```js
const els = document.querySelectorAll('#dropdown option:not([disabled])');
assert.isAtLeast(els.length, 2);
```

`#dropdown`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #dropdown');
assert.isNotNull(el);
```

`input` 요소가 최소 두 개 이상 있어야 하며, `type`가 `radio`인 라디오 버튼이어야 합니다.

```js
const els = document.querySelectorAll('input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

최소 두 개 이상의 라디오 버튼이 `#survey-form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#survey-form input[type="radio"]');
assert.isAtLeast(els.length, 2);
```

모든 라디오 버튼에 `value` 속성과 값이 있어야 합니다.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][value=""], input[type="radio"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

모든 라디오 버튼에 `name` 속성과 값이 있어야 합니다.

```js
const els1 = document.querySelectorAll('input[type="radio"]');
const els2 = document.querySelectorAll(
  'input[type="radio"][name=""], input[type="radio"]:not([name])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

모든 라디오 버튼 그룹에는 최소 2개의 라디오 버튼이 있어야 합니다.

```js
const radioButtons = document.querySelectorAll('input[type="radio"]');
const groups = {};

if (radioButtons) {
  radioButtons.forEach(el => {
    if (!groups[el.name]) groups[el.name] = [];
    groups[el.name].push(el);
  });
}

const groupKeys = Object.keys(groups);

groupKeys.forEach(key => {
  if (groups[key].length < 2) assert.fail();
});

assert.isAbove(groupKeys.length, 0);
```

`input` 요소가 최소 두 개 이상 있어야 하며, `type`가 `checkbox`인 체크박스여야 하며 `#survey-form`의 자손 요소여야 합니다.

```js
const els = document.querySelectorAll('#survey-form input[type="checkbox"]');
assert.isAtLeast(els.length, 2);
```

`#survey-form` 내 모든 체크박스에 `value` 속성과 값이 있어야 합니다.

```js
const els1 = document.querySelectorAll('#survey-form input[type="checkbox"]');
const els2 = document.querySelectorAll(
  '#survey-form input[type="checkbox"][value=""], #survey-form input[type="checkbox"]:not([value])'
);
assert.isAbove(els1.length, 0);
assert.lengthOf(els2, 0);
```

`textarea` 요소가 최소 하나 이상 있어야 하며 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form textarea');
assert.isNotNull(el);
```

`input` 또는 `button` 요소에 `id`가 `submit`인 속성이 있어야 합니다.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.isTrue(el.tagName === 'INPUT' || el.tagName === 'BUTTON');
```

`#submit`에 `type`가 `submit`인 속성이 있어야 합니다.

```js
const el = document.getElementById('submit');
assert.isNotNull(el);
assert.strictEqual(el.type, 'submit');
```

`#submit`는 `#survey-form`의 자손 요소여야 합니다.

```js
const el = document.querySelector('#survey-form #submit');
assert.isNotNull(el);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Survey Form</title>
  </head>

  <body></body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html>
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="container">
      <header class="header">
        <h1 id="title" class="text-center">freeCodeCamp Survey Form</h1>
        <p id="description" class="description text-center">
          Thank you for taking the time to help us improve the platform
        </p>
      </header>
      <form id="survey-form">
        <div class="form-group">
          <label id="name-label" for="name">Name</label>
          <input
            type="text"
            name="name"
            id="name"
            class="form-control"
            placeholder="Enter your name"
            required
          />
        </div>
        <div class="form-group">
          <label id="email-label" for="email">Email</label>
          <input
            type="email"
            name="email"
            id="email"
            class="form-control"
            placeholder="Enter your Email"
            required
          />
        </div>
        <div class="form-group">
          <label id="number-label" for="number"
            >Age<span class="clue">(optional)</span></label
          >
          <input
            type="number"
            name="age"
            id="number"
            min="10"
            max="99"
            class="form-control"
            placeholder="Age"
          />
        </div>
        <div class="form-group">
          <p>Which option best describes your current role?</p>
          <select id="dropdown" name="role" class="form-control" required>
            <option disabled selected value>Select current role</option>
            <option value="student">Student</option>
            <option value="job">Full Time Job</option>
            <option value="learner">Full Time Learner</option>
            <option value="preferNo">Prefer not to say</option>
            <option value="other">Other</option>
          </select>
        </div>

        <div class="form-group">
          <p>Would you recommend freeCodeCamp to a friend?</p>
          <label>
            <input
              name="user-recommend"
              value="definitely"
              type="radio"
              class="input-radio"
              checked
            />Definitely</label
          >
          <label>
            <input
              name="user-recommend"
              value="maybe"
              type="radio"
              class="input-radio"
            />Maybe</label
          >

          <label
            ><input
              name="user-recommend"
              value="not-sure"
              type="radio"
              class="input-radio"
            />Not sure</label
          >
        </div>

        <div class="form-group">
          <p>What is your favorite feature of freeCodeCamp?</p>
          <select id="most-like" name="mostLike" class="form-control" required>
            <option disabled selected value>Select an option</option>
            <option value="challenges">Challenges</option>
            <option value="projects">Projects</option>
            <option value="community">Community</option>
            <option value="openSource">Open Source</option>
          </select>
        </div>

        <div class="form-group">
          <p>
            What would you like to see improved?
            <span class="clue">(Check all that apply)</span>
          </p>

          <label
            ><input
              name="prefer"
              value="front-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Front-end Projects</label
          >
          <label>
            <input
              name="prefer"
              value="back-end-projects"
              type="checkbox"
              class="input-checkbox"
            />Back-end Projects</label
          >
          <label
            ><input
              name="prefer"
              value="data-visualization"
              type="checkbox"
              class="input-checkbox"
            />Data Visualization</label
          >
          <label
            ><input
              name="prefer"
              value="challenges"
              type="checkbox"
              class="input-checkbox"
            />Challenges</label
          >
          <label
            ><input
              name="prefer"
              value="open-source-community"
              type="checkbox"
              class="input-checkbox"
            />Open Source Community</label
          >
          <label
            ><input
              name="prefer"
              value="gitter-help-rooms"
              type="checkbox"
              class="input-checkbox"
            />Gitter help rooms</label
          >
          <label
            ><input
              name="prefer"
              value="videos"
              type="checkbox"
              class="input-checkbox"
            />Videos</label
          >
          <label
            ><input
              name="prefer"
              value="city-meetups"
              type="checkbox"
              class="input-checkbox"
            />City Meetups</label
          >
          <label
            ><input
              name="prefer"
              value="wiki"
              type="checkbox"
              class="input-checkbox"
            />Wiki</label
          >
          <label
            ><input
              name="prefer"
              value="forum"
              type="checkbox"
              class="input-checkbox"
            />Forum</label
          >
          <label
            ><input
              name="prefer"
              value="additional-courses"
              type="checkbox"
              class="input-checkbox"
            />Additional Courses</label
          >
        </div>

        <div class="form-group">
          <p>Any comments or suggestions?</p>
          <textarea
            id="comments"
            class="input-textarea"
            name="comment"
            placeholder="Enter your comment here..."
          ></textarea>
        </div>

        <div class="form-group">
          <button type="submit" id="submit" class="submit-button">
            Submit
          </button>
        </div>
      </form>
    </div>
  </body>
</html>
```
