---
id: 587d7fae367417b2b2512be3
title: الحصول على JSON باستخدام طريقة JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301502
dashedName: get-json-with-the-javascript-xmlhttprequest-method
---

# --description--

يمكنك أيضًا طلب بيانات من مصدر خارجي. هنا يأتي دور واجهات برمجة التطبيقات (APIs).

تذكر أن واجهات برمجة التطبيقات - أو واجهات برمجة التطبيقات للتطبيقات - هي أدوات تستخدمها الحواسيب للتواصل مع بعضها. ستتعلم كيف تحدّث HTML بالبيانات التي نحصل عليها من APIs باستخدام تقنية تُسمى AJAX.

معظم APIs الخاصة بالويب تنقل البيانات بصيغة تُسمى JSON. JSON تعني JavaScript Object Notation.

صيغة JSON تشبه جدًا صيغة كائنات JavaScript الحرفية. تحتوي JSON على خصائص الكائن وقيمها الحالية، محاطة بين `{` و `}`.

غالبًا ما تُسمى هذه الخصائص وقيمها "أزواج المفتاح-القيمة".

مع ذلك، JSON التي تنقلها APIs تُرسل كـ `bytes`، ويستقبلها تطبيقك كسلسلة نصية `string`. يمكن تحويلها إلى كائنات JavaScript، لكنها ليست كائنات JavaScript افتراضيًا. طريقة `JSON.parse` تُحلل السلسلة النصية وتُنشئ كائن JavaScript الموصوف بها.

يمكنك طلب JSON من API صور القطط الخاص بـ freeCodeCamp. إليك الشفرة التي يمكنك وضعها في حدث النقر للقيام بذلك:

```js
const req = new XMLHttpRequest();
req.open('GET', '/json/cats.json', true);
req.send();
req.onload = function () {
  const json = JSON.parse(req.responseText);
  document.getElementsByClassName('message')[0].innerHTML =
    JSON.stringify(json);
};
```

إليك مراجعة لما يقوم به كل جزء. كائن JavaScript `XMLHttpRequest` يحتوي على عدد من الخصائص والطرق التي تُستخدم لنقل البيانات. أولًا، يُنشأ مثيل من كائن `XMLHttpRequest` ويُحفظ في المتغير `req`. بعد ذلك، تُهيئ طريقة `open` الطلب - في هذا المثال يتم طلب بيانات من API، لذلك هو طلب `GET`. المعلمة الثانية لـ `open` هي عنوان URL الخاص بالـ API الذي تطلب منه البيانات. المعلمة الثالثة هي قيمة منطقية حيث يجعل `true` الطلب غير متزامن. طريقة `send` ترسل الطلب. أخيرًا، معالِج حدث `onload` يُحلل البيانات المُعادة ويطبق طريقة `JSON.stringify` لتحويل كائن JavaScript إلى سلسلة نصية. تُدرج هذه السلسلة كنص الرسالة.

# --instructions--

حدّث الشفرة لإنشاء وإرسال طلب `GET` إلى API صور القطط الخاص بـ freeCodeCamp. ثم انقر زر `Get Message`. ستستبدل دالة AJAX النص `The message will go here` بالإخراج الخام لـ JSON من الـ API.

# --hints--

يجب أن تنشئ شفرتك كائن `XMLHttpRequest` جديدًا.

```js
assert.match(code, /new\s+?XMLHttpRequest\(\s*?\)/g);
```

يجب أن تستخدم شفرتك طريقة `open` لتهيئة طلب `GET` إلى API صور القطط الخاص بـ freeCodeCamp.

```js
assert.match(
  code,
  /\.open\(\s*('|")GET\1\s*,\s*('|")\/json\/cats\.json\2\s*,\s*true\s*\)/g
);
```

يجب أن تستخدم شفرتك طريقة `send` لإرسال الطلب.

```js
assert.match(code, /\.send\(\s*\)/g);
```

يجب أن تحتوي شفرتك على معالِج حدث `onload` مضبوط على دالة.

```js
assert.match(
  code,
  /\.onload\s*=\s*(function|\(\s*\))\s*(\(\s*\)|\=\>)\s*{/g
);
```

يجب أن تستخدم شفرتك طريقة `JSON.parse` لتحليل `responseText`.

```js
assert.match(code, /JSON\s*\.parse\(\s*.*\.responseText\s*\)/g);
```

يجب أن تحصل شفرتك على العنصر الذي يحمل الصنف `message` وتغير محتوى HTML الداخلي إلى سلسلة بيانات JSON.

```js
assert.match(
  code,
  /document\s*\.getElementsByClassName\(\s*?('|")message\1\s*?\)\[0\]\s*\.innerHTML\s*?=\s*?JSON\.stringify\(.+?\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = () => {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
