---
id: 587d7fae367417b2b2512be4
title: الوصول إلى بيانات JSON من واجهة برمجة التطبيقات
challengeType: 6
forumTopicId: 301499
dashedName: access-the-json-data-from-an-api
---

# --description--

في التحدٍّ السابق، رأيت كيف تحصل على بيانات JSON من واجهة برمجة التطبيقات لصور القطط في freeCodeCamp.

الآن ستلقي نظرة أقرب على البيانات المُعادة لفهم تنسيق JSON بشكل أفضل. تذكر بعض الترميزات في JavaScript:

<blockquote>[ ] -> تمثل الأقواس المربعة مصفوفة.<br>{ } -> تمثل الأقواس المعقوفة كائنًا.<br>" " -> تمثل علامات الاقتباس المزدوجة سلسلة نصية. تُستخدم أيضًا لأسماء المفاتيح في JSON.</blockquote>

فهم بنية البيانات التي تُعيدها واجهة برمجة التطبيقات مهم لأنه يؤثر على كيفية استرجاع القيم التي تحتاجها.

على اليمين، انقر على زر `Get Message` لتحميل JSON الخاص بواجهة برمجة التطبيقات لصور القطط في freeCodeCamp إلى HTML.

أول وآخر حرف تراه في بيانات JSON هما أقواس مربعة `[ ]`. هذا يعني أن البيانات المُعادة هي مصفوفة. الحرف الثاني في بيانات JSON هو قوس معقوف `{`، الذي يبدأ كائنًا. بالنظر عن كثب، ترى أن هناك ثلاثة كائنات منفصلة. بيانات JSON هي مصفوفة من ثلاثة كائنات، يحتوي كل منها على معلومات عن صورة قطة.

تعلمت سابقًا أن الكائنات تحتوي على "أزواج المفتاح-القيمة" التي تفصل بينها فواصل. في مثال صورة القطة، يحتوي الكائن الأول على `"id":0` حيث `id` هو المفتاح و`0` هي القيمة المقابلة له. بالمثل، هناك مفاتيح لـ `imageLink` و`altText` و`codeNames`. كل كائن صورة قطة له هذه المفاتيح نفسها، لكن بقيم مختلفة.

زوج مفتاح-قيمة آخر مثير للاهتمام في الكائن الأول هو `"codeNames":["Juggernaut","Mrs. Wallace","ButterCup"]`. هنا `codeNames` هو المفتاح وقيمته مصفوفة من ثلاث سلاسل نصية. من الممكن أن تكون هناك مصفوفات من كائنات، وكذلك مفتاح بقيمة مصفوفة.

تذكر كيف تصل إلى البيانات في المصفوفات والكائنات. تستخدم المصفوفات ترميز الأقواس للوصول إلى فهرس محدد من عنصر. تستخدم الكائنات إما ترميز الأقواس أو ترميز النقطة للوصول إلى قيمة خاصية معينة. إليك مثالًا يطبع خاصية `altText` لأول صورة قطة - لاحظ أن بيانات JSON المحللة في المحرر محفوظة في متغير يسمى `json`:

```js
console.log(json[0].altText);
```

ستعرض وحدة التحكم السلسلة النصية `A white cat wearing a green helmet shaped melon on its head.`.

# --instructions--

بالنسبة للقطة التي لها `id` يساوي 2، اطبع إلى وحدة التحكم القيمة الثانية في مصفوفة `codeNames`. يجب أن تستخدم ترميز الأقواس وترميز النقطة على الكائن (المحفوظ في المتغير `json`) للوصول إلى القيمة.

# --hints--

يجب أن يستخدم شفرتك ترميز الأقواس وترميز النقطة للوصول إلى اسم الكود الصحيح، واطبع `Loki` إلى وحدة التحكم.

```js
assert.match(
  code,
  /console\s*\.\s*log\s*\(\s*json\s*\[2\]\s*(\.\s*codeNames|\[\s*('|`|")codeNames\2\s*\])\s*\[\s*1\s*\]\s*\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line

        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        document.getElementsByClassName('message')[0].innerHTML =
          JSON.stringify(json);
        // Add your code below this line
        console.log(json[2].codeNames[1]);
        // Add your code above this line
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
