---
id: 587d7fad367417b2b2512be1
title: التعامل مع أحداث النقر باستخدام JavaScript عبر الخاصية onclick
challengeType: 6
forumTopicId: 301503
dashedName: handle-click-events-with-javascript-using-the-onclick-property
---

# --description--

تريد أن تُنفذ الشفرة فقط بعد انتهاء تحميل الصفحة. لهذا الغرض، يمكنك ربط حدث JavaScript بالوثيقة يسمى `DOMContentLoaded`. إليك الشفرة التي تقوم بذلك:

```js
document.addEventListener('DOMContentLoaded', function () {});
```

يمكنك تنفيذ معالجات الأحداث داخل دالة `DOMContentLoaded`. يمكنك تنفيذ معالج حدث `onclick` الذي يُفعّل عند نقر المستخدم على العنصر `#getMessage`، وذلك بإضافة الشفرة التالية:

```js
document.getElementById('getMessage').onclick = function () {};
```

# --instructions--

أضف معالج حدث نقر داخل دالة `DOMContentLoaded` للعنصر الذي يحمل المعرف `getMessage`.

# --hints--

يجب أن تستخدم الشفرة طريقة `document.getElementById` لاختيار العنصر الذي معرفه `getMessage`.

```js
assert.match(code, /document\s*\.getElementById\(\s*('|")getMessage\1\s*\)/g);
```

يجب أن تضيف الشفرة معالج حدث `onclick`.

```js
assert.isFunction(document.getElementById('getMessage').onclick);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    // Add your code below this line
    document.getElementById('getMessage').onclick = function () {};
    // Add your code above this line
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
