---
id: 587d7faf367417b2b2512be9
title: إرسال البيانات باستخدام طريقة JavaScript XMLHttpRequest
challengeType: 6
forumTopicId: 301504
dashedName: post-data-with-the-javascript-xmlhttprequest-method
---

# --description--

في الأمثلة السابقة، استلمت بيانات من مورد خارجي. يمكنك أيضًا إرسال بيانات إلى مورد خارجي، طالما أن ذلك المورد يدعم طلبات AJAX وتعرف عنوان URL الخاص به.

تُستخدم طريقة JavaScript `XMLHttpRequest` أيضًا لإرسال البيانات إلى خادم. إليك مثالًا:

```js
const xhr = new XMLHttpRequest();
xhr.open('POST', url, true);
xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
xhr.onreadystatechange = function () {
  if (xhr.readyState === 4 && xhr.status === 201) {
    const serverResponse = JSON.parse(xhr.response);
    document.getElementsByClassName('message')[0].textContent =
      serverResponse.userName + serverResponse.suffix;
  }
};
const body = JSON.stringify({ userName: userName, suffix: ' loves cats!' });
xhr.send(body);
```

لقد رأيت عدة من هذه الطرق سابقًا. هنا تُهيئ الطريقة `open` الطلب كـ `POST` إلى عنوان URL المعطى للمورد الخارجي، وتمرر `true` كمعلمة ثالثة - مما يشير إلى تنفيذ العملية بشكل غير متزامن.

تُعيّن الطريقة `setRequestHeader` قيمة قسم علوي لطلب HTTP، يحتوي على معلومات عن المرسل والطلب. يجب استدعاؤها بعد الطريقة `open`، ولكن قبل الطريقة `send`. المعلمتان هما اسم القسم العلوي والقيمة التي تُعيّن كمحتوى لذلك القسم.

بعد ذلك، يستمع معالِج الحدث `onreadystatechange` لتغير حالة الطلب. تعني قيمة `readyState` التي تساوي `4` أن العملية اكتملت، و`status` التي تساوي `201` تعني أن الطلب نجح. لذلك، يمكن تحديث شفرة HTML للوثيقة.

أخيرًا، ترسل الطريقة `send` الطلب مع قيمة `body`. يتكون `body` من مفتاحي `userName` و`suffix`.

# --instructions--

حدّث الشفرة لتجعلها ترسل طلب `POST` إلى نقطة النهاية الخاصة بواجهة برمجة التطبيقات. ثم اكتب اسمك في حقل الإدخال وانقر على `Send Message`. يجب أن تستبدل دالة AJAX النص `Reply from Server will be here.` بالبيانات القادمة من الخادم. نسق الاستجابة لعرض اسمك مضافًا إليه النص ` loves cats`.

# --hints--

يجب أن تنشئ شفرتك كائن `XMLHttpRequest` جديدًا.

```js
assert.match(code, /new\s*XMLHttpRequest\(\s*\)/g);
```

يجب أن تستخدم شفرتك الطريقة `open` لتهيئة طلب `POST` إلى الخادم.

```js
assert.match(code, /\.open\(\s*('|")POST\1\s*,\s*url\s*,\s*true\s*\)/g);
```

يجب أن تستخدم شفرتك الطريقة `setRequestHeader`.

```js
assert.match(
  code,
  /\.setRequestHeader\(\s*('|")Content-Type\1\s*,\s*('|")application\/json;\s*charset=UTF-8\2\s*\)/g
);
```

يجب أن تحتوي شفرتك على معالِج حدث `onreadystatechange` مضبوط على دالة.

```js
assert.match(code, /\.onreadystatechange\s*=/g);
```

يجب أن تحصل شفرتك على العنصر الذي يحمل الصنف `message` وتغير خاصية `textContent` إلى `userName loves cats`

```js
assert.match(
  code,
  /document\.getElementsByClassName\(\s*('|")message\1\s*\)\[0\]\.textContent\s*=\s*.+?\.userName\s*\+\s*.+?\.suffix/g
);
```

يجب أن تستخدم شفرتك الطريقة `send`.

```js
assert.match(code, /\.send\(\s*body\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line

      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message box">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('sendMessage').onclick = function () {
      const userName = document.getElementById('name').value;
      const url = 'https://jsonplaceholder.typicode.com/posts';
      // Add your code below this line
      const xhr = new XMLHttpRequest();
      xhr.open('POST', url, true);
      xhr.setRequestHeader('Content-Type', 'application/json; charset=UTF-8');
      xhr.onreadystatechange = function () {
        if (xhr.readyState === 4 && xhr.status === 201) {
          const serverResponse = JSON.parse(xhr.response);
          document.getElementsByClassName('message')[0].textContent =
            serverResponse.userName + serverResponse.suffix;
        }
      };
      const body = JSON.stringify({
        userName: userName,
        suffix: ' loves cats!'
      });
      xhr.send(body);
      // Add your code above this line
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Friends</h1>
<p class="message">Reply from Server will be here</p>
<p>
  <label for="name"
    >Your name:
    <input type="text" id="name" />
  </label>
  <button id="sendMessage">Send Message</button>
</p>
```
