---
id: 587d7faf367417b2b2512be8
title: الحصول على بيانات الموقع الجغرافي لإيجاد إحداثيات GPS للمستخدم
challengeType: 6
forumTopicId: 18188
dashedName: get-geolocation-data-to-find-a-users-gps-coordinates
---

# --description--

شيء رائع آخر يمكنك فعله هو الوصول إلى الموقع الحالي للمستخدم. كل متصفح يحتوي على كائن navigator مدمج يمكنه تزويدك بهذه المعلومات.

سيحصل navigator على خط الطول وخط العرض الحاليين للمستخدم.

ستظهر لك رسالة تطلب السماح أو الحظر لهذا الموقع لمعرفة موقعك الحالي. يمكن إكمال التحدي بأي من الطريقتين، طالما أن الكود صحيح.

باختيار السماح، سترى النص على الهاتف الناتج يتغير إلى خط العرض وخط الطول الخاصين بك.

إليك كود يقوم بذلك:

```js
if (navigator.geolocation) {
  navigator.geolocation.getCurrentPosition(function (position) {
    document.getElementById('data').innerHTML =
      'latitude: ' +
      position.coords.latitude +
      '<br>longitude: ' +
      position.coords.longitude;
  });
}
```

أولًا، يتحقق مما إذا كان كائن `navigator.geolocation` موجودًا. إذا كان موجودًا، تُستدعى الطريقة `getCurrentPosition` على هذا الكائن، والتي تبدأ طلبًا غير متزامن لموقع المستخدم. إذا نجح الطلب، تُنفذ دالة الاستدعاء الراجع في الطريقة. تصل هذه الدالة إلى قيم خط العرض وخط الطول في كائن `position` باستخدام ترميز النقطة وتُحدّث الـ HTML.

# --instructions--

أضف كود المثال داخل وسم `script` للتحقق من موقع المستخدم الحالي وإدخاله في HTML.

# --hints--

يجب أن يستخدم كودك `navigator.geolocation` للوصول إلى الموقع الحالي للمستخدم.

```js
assert.match(code, /navigator\.geolocation\.getCurrentPosition/g);
```

يجب أن يستخدم كودك `position.coords.latitude` لعرض الموقع العرضي للمستخدم.

```js
assert.match(code, /position\.coords\.latitude/g);
```

يجب أن يستخدم كودك `position.coords.longitude` لعرض الموقع الطولي للمستخدم.

```js
assert.match(code, /position\.coords\.longitude/g);
```

يجب أن تعرض موقع المستخدم داخل عنصر `div` الذي يحمل `id="data"`.

```js
assert.match(
  code,
  /document\.getElementById\(\s*('|")data\1\s*\)\.innerHTML/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  // Add your code below this line

  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data"></div>
```

# --solutions--

```html
<script>
  // Add your code below this line
  if (navigator.geolocation) {
    navigator.geolocation.getCurrentPosition(function(position) {
      document.getElementById('data').innerHTML = "latitude: " + position.coords.latitude + "<br>longitude: " + position.coords.longitude;
    });
  }
  // Add your code above this line
</script>
<h4>You are here:</h4>
<div id="data">

</div>

</section>
```
