---
id: 587d7fae367417b2b2512be5
title: تحويل بيانات JSON إلى HTML
challengeType: 6
forumTopicId: 16807
dashedName: convert-json-data-to-html
---

# --description--

الآن بعد أن تحصل على البيانات من واجهة برمجة تطبيقات JSON، يمكنك عرضها في HTML.

يمكنك استخدام طريقة `forEach` للتكرار عبر البيانات لأن كائنات صور القطط محفوظة في مصفوفة. عند الوصول إلى كل عنصر، يمكنك تعديل عناصر HTML.

أولًا، أعلن عن متغير html باستخدام `let html = "";`.

ثم، كرر عبر JSON، مضيفًا HTML إلى المتغير يلف أسماء المفاتيح داخل وسم `strong`، متبوعًا بالقيمة. عند الانتهاء من التكرار، اعرضه.

إليك الشفرة التي تقوم بذلك:

```js
let html = '';
json.forEach(function (val) {
  const keys = Object.keys(val);
  html += "<div class = 'cat'>";
  keys.forEach(function (key) {
    html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
  });
  html += '</div><br>';
});
```

**ملاحظة:** لهذا التحدٍّ، تحتاج إلى إضافة عناصر HTML جديدة إلى الصفحة، لذا لا يمكنك الاعتماد على `textContent`. بدلًا من ذلك، عليك استخدام `innerHTML`، الذي قد يجعل الموقع عرضة لهجمات البرمجة النصية عبر المواقع.

# --instructions--

أضف طريقة `forEach` للتكرار عبر بيانات JSON وإنشاء عناصر HTML لعرضها.

إليك بعض بيانات JSON كمثال:

```json
[
  {
    "id": 0,
    "imageLink": "https://cdn.freecodecamp.org/curriculum/legacy-json-apis-ajax/funny-cat.jpg",
    "altText": "A white cat wearing a green helmet shaped melon on its head. ",
    "codeNames": ["Juggernaut", "Mrs. Wallace", "Buttercup"]
  }
]
```

# --hints--

يجب أن يخزن شفرتك البيانات في المتغير `html`

```js
assert.match(__helpers.removeWhiteSpace(code), /html(\+=|=html\+)/g);
```

يجب أن تستخدم شفرتك طريقة `forEach` للتكرار عبر بيانات JSON من واجهة برمجة التطبيقات.

```js
assert.match(code, /json\.forEach/g);
```

يجب أن تلف أسماء المفاتيح في شفرتك داخل وسم `strong`.

```js
assert.match(code, /<strong>.+<\/strong>/g);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line

        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = function () {
      const req = new XMLHttpRequest();
      req.open('GET', '/json/cats.json', true);
      req.send();
      req.onload = function () {
        const json = JSON.parse(req.responseText);
        let html = '';
        // Add your code below this line
        json.forEach(function (val) {
          var keys = Object.keys(val);
          html += "<div class = 'cat'>";
          keys.forEach(function (key) {
            html += '<strong>' + key + '</strong>: ' + val[key] + '<br>';
          });
          html += '</div><br>';
        });
        // Add your code above this line
        document.getElementsByClassName('message')[0].innerHTML = html;
      };
    };
  });
</script>

<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>

<h1>Cat Photo Finder</h1>
<p class="message">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
