---
id: 5ccfad82bb2dc6c965a848e5
title: الحصول على JSON باستخدام طريقة fetch في JavaScript
challengeType: 6
forumTopicId: 301501
dashedName: get-json-with-the-javascript-fetch-method
---

# --description--

طريقة أخرى لطلب بيانات خارجية هي استخدام طريقة `fetch()` في JavaScript. تعادل `XMLHttpRequest`، لكن الصيغة النحوية تُعتبر أسهل للفهم.

إليك الشفرة لعمل طلب GET إلى `/json/cats.json`

```js
fetch('/json/cats.json')
  .then(response => response.json())
  .then(data => {
    document.getElementById('message').innerHTML = JSON.stringify(data);
  });
```

ملاحظة: تستخدم طريقة `fetch()` طريقة `GET` كطريقة `HTTP` افتراضية. هذا يعني أنك لا تحتاج إلى تحديدها صراحةً لاسترجاع البيانات الأساسية.

ألقِ نظرة على كل جزء من هذه الشفرة.

السطر الأول هو الذي ينفذ الطلب. إذًا، `fetch(URL)` ينفذ طلب `GET` إلى عنوان URL المحدد. الطريقة تُرجع وعدًا.

بعد إرجاع الوعد، إذا نجح الطلب، تُنفذ طريقة `then` التي تأخذ الاستجابة وتحولها إلى صيغة JSON.

طريقة `then` تُرجع أيضًا وعدًا، يُعالج بواسطة طريقة `then` التالية. المعلمة في `then` الثانية هي كائن JSON الذي تبحث عنه!

الآن، يختار العنصر الذي سيستقبل البيانات باستخدام `document.getElementById()`. ثم يُعدل شفرة HTML للعنصر بإدخال سلسلة نصية مُنشأة من كائن JSON المُرجع من الطلب.

# --instructions--

حدّث الشفرة لإنشاء وإرسال طلب `GET` إلى واجهة برمجة تطبيقات صور القطط في freeCodeCamp. لكن هذه المرة، باستخدام طريقة `fetch` بدلًا من `XMLHttpRequest`.

# --hints--

يجب أن تستخدم شفرتك البيانات المسترجعة لاستبدال محتوى HTML الداخلي

```js
const catData = 'dummy data';
const ref = fetch;
fetch = () => Promise.resolve({ json: () => catData });

  try {
    document.getElementById('getMessage').click();
    await new Promise((resolve, reject) => setTimeout(() => resolve(), 250));
  } catch (error) {
    console.log(error);
  } finally {
    fetch = ref;
    assert.equal(
      document.getElementById('message').textContent,
      JSON.stringify(catData)
    );
  }

```

يجب أن تنفذ شفرتك طلب `GET` باستخدام `fetch`.

```js
assert.match(code, /fetch\s*\(\s*('|")\/json\/cats\.json\1\s*\)/g);
```

يجب أن تستخدم شفرتك `then` لتحويل الاستجابة إلى JSON.

```js
assert.match(
  code,
  /\.then\s*\(\s*\(?(?<var>\w+)\)?\s*=>\s*\k<var>\s*\.json\s*\(\s*\)\s*\)/g
);
```

يجب أن تستخدم شفرتك `then` لمعالجة البيانات المحولة إلى JSON بواسطة `then` الآخر.

```js
assert.match(__helpers.removeWhiteSpace(code), /\.then\(\(?\w+\)?=>{[^}]*}\)/g);
```

يجب أن تحصل شفرتك على العنصر الذي يحمل المعرف `message` وتغير محتوى HTML الداخلي إلى سلسلة بيانات JSON.

```js
assert.match(
  __helpers.removeWhiteSpace(code),
  /document\.getElementById\(('|")message\1\)\.innerHTML=JSON\.stringify\(?\w+\)/g
);
```

# --seed--

## --seed-contents--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      // Add your code below this line
      // Add your code above this line
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```

# --solutions--

```html
<script>
  document.addEventListener('DOMContentLoaded', function () {
    document.getElementById('getMessage').onclick = () => {
      fetch('/json/cats.json')
        .then(response => response.json())
        .then(data => {
          document.getElementById('message').innerHTML = JSON.stringify(data);
        });
    };
  });
</script>
<style>
  body {
    text-align: center;
    font-family: 'Helvetica', sans-serif;
  }
  h1 {
    font-size: 2em;
    font-weight: bold;
  }
  .box {
    border-radius: 5px;
    background-color: #eee;
    padding: 20px 5px;
  }
  button {
    color: white;
    background-color: #4791d0;
    border-radius: 5px;
    border: 1px solid #4791d0;
    padding: 5px 10px 8px 10px;
  }
  button:hover {
    background-color: #0f5897;
    border: 1px solid #0f5897;
  }
</style>
<h1>Cat Photo Finder</h1>
<p id="message" class="box">The message will go here</p>
<p>
  <button id="getMessage">Get Message</button>
</p>
```
