---
id: 6995d0325fd7da40854fac40
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У попередньому уроці ви дізналися про бінарний пошук і чим він відрізняється від лінійного пошуку.

У цьому практичному занятті ви застосуєте це на практиці, створивши алгоритм бінарного пошуку.

Для початку створіть функцію `binarySearch` з параметрами `searchList` та `value`.

`searchList` — це масив елементів, у якому функція буде шукати, а `value` — це ціль, яку функція намагатиметься знайти в `searchList`.

# --hints--

Ви повинні створити функцію `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Ваша функція `binarySearch` має мати параметри `searchList` та `value`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
assert.lengthOf(binarySearch?.parameters, 2);
assert.equal(binarySearch?.parameters[0].toString(), "searchList");
assert.equal(binarySearch?.parameters[1].toString(), "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
