---
id: 6995d0325fd7da40854fac40
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У попередньому уроці ви дізналися про бінарний пошук і чим він відрізняється від лінійного пошуку.

У цьому практичному занятті ви застосуєте це на практиці, створивши алгоритм бінарного пошуку.

Для початку створіть функцію `binarySearch` з параметрами `searchList` та `value`.

`searchList` — це масив елементів, у якому функція буде шукати, а `value` — це ціль, яку функція намагатиметься знайти в `searchList`.

# --hints--

Ви повинні створити функцію `binarySearch`.

```js
assert.isFunction(binarySearch);
```

Ваша функція `binarySearch` має мати параметри `searchList` та `value`.

```js
const binSearchParams = __helpers.getFunctionParams(binarySearch.toString());

const binarySearchParamsNames = binSearchParams.map(param => param.name)
assert.include(binarySearchParamsNames, "searchList");
assert.include(binarySearchParamsNames, "value");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
