---
id: 6996ca21a74104721f2c7e5d
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Вам потрібно протестувати деякі речі, щоб зрозуміти потік алгоритму на цьому ініціальному етапі.

Для цього спочатку потрібно вийти з циклу. Це тому, що поточна реалізація дозволить лише одну ітерацію, тож якщо умова не виконується, виникне нескінченний цикл.

Одразу після інструкції `if` використайте ключове слово `break`, щоб вийти з циклу `while`. Потім, після циклу `while`, поверніть порожній масив, щоб позначити, що значення не було знайдено.

# --hints--

Ви повинні використати ключове слово `break`, щоб вийти з циклу `while` після інструкції `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?\}\s*break\s*;?/
);
```

Після вашого циклу `while` слід повернути порожній масив.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /return\s*\[\s*\]/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    if (value === valueAtMiddle) {
      return pathToTarget;
    }
--fcc-editable-region--
    
  }
  
--fcc-editable-region--
}
```
