---
id: 6996c7363a4a437839f1f15d
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Тепер, коли ви визначили межі свого пошуку, настав час створити цикл, який виконуватиме алгоритм бінарного пошуку.

Бінарний пошук працює шляхом багаторазового звуження простору пошуку. Цей процес триває доти, доки існує дійсний діапазон елементів для перевірки.

Ви можете виразити це за допомогою `while` циклу, який триватиме, доки ваш `low` вказівник менший або дорівнює `high` вказівнику.

Якщо `low` коли-небудь стане більшим за `high`, це означає, що простір пошуку став порожнім, і `value` немає в масиві.

Всередині функції `binarySearch` створіть `while` цикл з умовою, яка перевіряє, чи `low` менший або дорівнює high.

# --hints--

Ви повинні створити `while` цикл, який працюватиме, поки `low` менший або дорівнює `high`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /while\s*\(\s*low\s*<=\s*high\s*\)/);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
