---
id: 6996cb5869acf58de751ae81
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Нагадаємо, що функція приймає масив під назвою `searchList` та `value`, який ви шукаєте.

Тепер викличте функцію з `binarySearch([1, 2, 3, 4, 5], 3)` і `binarySearch([1, 2, 3, 4, 5, 9], 4)` та виведіть обидві інструкції в консоль.

# --hints--

Ви повинні викликати функцію з `binarySearch([1, 2, 3, 4, 5], 3)` і вивести це в консоль.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5],3))";
assert.include(strippedCode, expectedStatement);
```

Ви повинні викликати функцію з `binarySearch([1, 2, 3, 4, 5, 9], 4)` і вивести це в консоль.

```js
const strippedCode = __helpers.removeWhiteSpace(__helpers.removeJSComments(code));
const expectedStatement = "console.log(binarySearch([1,2,3,4,5,9],4))";
assert.include(strippedCode, expectedStatement);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;

  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    }
    break;
  }
  return [];
}

--fcc-editable-region--

--fcc-editable-region--
```
