---
id: 6996c675b599baa7b103192d
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У бінарному пошуку існує найнижчий можливий індекс і найвищий можливий індекс.

Найнижчий можливий індекс представляє ліву межу в поточній частині значень, у якій виконується пошук, у цьому випадку `searchList`. Найвищий можливий індекс, навпаки, представляє праву межу в значеннях, у яких виконується пошук.

Щоб врахувати ці два, потрібно розглянути весь масив, у якому виконується пошук. Тож всередині функції `binarySearch` оголосіть змінну `low` зі значенням `0` і змінну `high` зі значенням, яке відповідає останньому індексу масиву, у якому виконується пошук.

# --hints--

Ви повинні оголосити змінну з ім’ям `low`, встановлену в `0`, у вашій функції `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { low } = binarySearch?.variables ?? {};
assert.isTrue(low?.value.matches("0"));
```

Ви повинні оголосити змінну з ім’ям `high`, встановлену в `searchList.length - 1`, у вашій функції `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { high } = binarySearch?.variables ?? {};
assert.isTrue(high?.value.matches("searchList.length - 1"));
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
