---
id: 6996c675b599baa7b103192d
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

У бінарному пошуку існує найнижчий можливий індекс і найвищий можливий індекс.

Найнижчий можливий індекс представляє ліву межу в поточній частині значень, у якій виконується пошук, у цьому випадку `searchList`. Найвищий можливий індекс, навпаки, представляє праву межу в значеннях, у яких виконується пошук.

Щоб врахувати ці дві межі, потрібно розглядати весь масив, у якому виконується пошук. Тож всередині функції `binarySearch` визначте змінну `low` зі значенням `0` і змінну `high` зі значенням, яке відповідає останньому індексу масиву, у якому виконується пошук.

# --hints--

У вашій функції `low` має бути змінна з ім’ям `0`, встановлена в `binarySearch`.

```js
assert.match(__helpers.removeJSComments(String(binarySearch)), /(const|let|var)\s+low\s*=\s*0/);
```

У вашій функції `high` має бути змінна з ім’ям `searchList.length - 1`, встановлена в `binarySearch`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+high\s*=\s*searchList\.length\s*-\s*1/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
--fcc-editable-region--
  
--fcc-editable-region--
}
```
