---
id: 6996c8375f4301004459ee85
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Тепер, коли ви обчислили індекс `mid`, потрібно фактично отримати значення з `searchList` за цим індексом. Це значення ви будете порівнювати з вашим параметром `value`, який є ціллю, яку ви шукаєте.

Створіть змінну `valueAtMiddle` і призначте їй елемент з `searchList`, що знаходиться за індексом `mid`.

# --hints--

Ви маєте створити змінну `valueAtMiddle`, призначивши їй `searchList[mid]`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+valueAtMiddle\s*=\s*searchList\s*\[\s*mid\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
