---
id: 6996cfb6c01e9a802add38be
title: Крок 13
challengeType: 1
dashedName: step-13
---

# --description--

Нарешті, якщо `value` не знайдено посередині і він також не більший за `valueAtMiddle`, тоді `value` має бути меншим за `valueAtMiddle` і розташованим зліва. Фінальний блок `else` обробить це.

Отже, завершіть цикл, видаливши ключове слово `break` і додавши блок `else`. Усередині нього оновіть змінну `high`, віднявши `1` від `mid`.

# --hints--

Ви повинні замінити ключове слово `break` на блок `else`, який оновить змінну `high` до `mid - 1`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s*\{\s*high\s*=\s*mid\s*-\s*1\s*;?\s*\}/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    }

--fcc-editable-region--      
    break;
--fcc-editable-region--

  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
