---
id: 6996d0a2c2ac55041145a13f
title: Крок 14
challengeType: 1
dashedName: step-14
---

# --description--

Ви помітили, що другий виклик функції тепер має значення `3, 5, 4` у списку? Це означає, що бінарний пошук працює так, як задумано.

Ось як це сталося: алгоритм спочатку перевірив `3` як середину початкового списку. Оскільки `4` більше за `3`, пошук змістився до правої половини. Потім він розглянув `5` як нову середину. Оскільки `4` менше за `5`, пошук перемістився вліво, врешті-решт визначивши `4` як середину кінцевого діапазону.

Щоб знову перевірити функцію, викличте її з `[1, 3, 5, 9, 14, 22], 10` і виведіть у консоль. Це ситуація, коли значення не буде знайдено.

# --hints--

Ви повинні викликати функцію з `binarySearch([1, 3, 5, 9, 14, 22], 10)` і вивести результат у консоль.

```js
assert.match(
  __helpers.removeJSComments(code),
  /console\s*\.\s*log\s*\(\s*binarySearch\s*\(\s*\[\s*1\s*,\s*3\s*,\s*5\s*,\s*9\s*,\s*14\s*,\s*22\s*\]\s*,\s*10\s*\)\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);

    if (value === valueAtMiddle) {
      return pathToTarget;
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
--fcc-editable-region--

--fcc-editable-region--
```
