---
id: 6996c60281feca6718cc15c1
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Усередині вашої `binarySearch` функції створіть змінну `pathToTarget` і ініціалізуйте її порожнім масивом для початку.

# --hints--

У вас має бути змінна `pathToTarget`, встановлена як порожній масив `[]` всередині вашої `binarySearch` функції.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+pathToTarget\s*=\s*\[\s*\]/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
