---
id: 6996c8f94fe6b250f11a2a1e
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер додайте `valueAtMiddle` до масиву `pathToTarget`, щоб відстежувати кроки, зроблені під час пошуку, незалежно від того, чи є це ціллю. Ви можете зробити це за допомогою методу `push()`.

# --hints--

Ви повинні використати метод `push()`, щоб додати `valueAtMiddle` до масиву `pathToTarget`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
