---
id: 6996ccadf9048c79b857bb15
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Ви бачите, що перший виклик повертає `[3]`. Це тому, що `3` — це середнє значення в `[1, 2, 3, 4, 5]`.

З іншого боку, другий виклик повертає порожній масив. Це відбувається тому, що після перевірки середнього значення цикл наразі не має логіки для звуження діапазону пошуку шляхом оновлення змінних `low` або `high`. Він лише виконує одну перевірку посередині і завершується без подальшого пошуку.

Щоб дозволити бінарному пошуку продовжувати звужувати свій пошук, додайте інструкцію `else if`, яка перевіряє, чи `value` більше за `valueAtMiddle`.

# --hints--

Ви повинні додати інструкцію `else if` з умовою `value > valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /else\s+if\s*\(\s*value\s*>\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    if (value === valueAtMiddle) {
      return pathToTarget;
    } 
--fcc-editable-region--
    break;
  }
  return [];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
```
