---
id: 6996c8f94fe6b250f11a2a1e
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Тепер додайте `valueAtMiddle` до масиву `pathToTarget`, щоб відстежувати кроки, зроблені під час пошуку, незалежно від того, чи є це ціллю. Ви можете зробити це за допомогою методу `push()`.

# --hints--

Ви повинні використати метод `push()`, щоб додати `valueAtMiddle` до масиву `pathToTarget`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearch?.toString());
assert.match(
  codeWithoutComments,
  /pathToTarget\s*\.\s*push\(\s*valueAtMiddle\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  const pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    const mid = Math.floor((low + high) / 2);
    const valueAtMiddle = searchList[mid];
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
