---
id: 6996c60281feca6718cc15c1
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Усередині вашої `binarySearch` функції створіть змінну `pathToTarget` і ініціалізуйте її порожнім масивом для початку.

# --hints--

Ви повинні оголосити змінну `pathToTarget`, ініціалізувавши її порожнім масивом, всередині вашої функції `binarySearch`.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch } = explorer.allFunctions;
const { pathToTarget } = binarySearch?.variables ?? {};
const value = eval(pathToTarget?.value.toString())
assert.isArray(value);
assert.isEmpty(value);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
