---
id: 6996d18b31aab5cbbffbc731
title: Крок 16
challengeType: 1
dashedName: step-16
---

# --description--

Останнє невелике покращення, яке ви можете додати, — це також показати індекс, на якому знаходиться ціль `value`.

Отже, замість того, щоб повертати лише `pathToTarget` в `if`-інструкції, повертайте масив, що містить `pathToTarget` і шаблонний рядок із повідомленням `Value found at index ${mid}`.

З цим ваше практичне заняття з алгоритму бінарного пошуку завершене!

# --hints--

Ви повинні оновити `return`-інструкцію в `if`, щоб функція повертала масив із `pathToTarget` як першим елементом і шаблонним рядком для `Value found at index ${mid}` як другим елементом. Обов’язково використовуйте зворотні лапки для шаблонного рядка.

```js
const explorer = await __helpers.Explorer(code);
const { binarySearch: binarySearchFn } = explorer.allFunctions;
const codeWithoutComments = __helpers.removeJSComments(binarySearchFn?.toString());
assert.match(
  codeWithoutComments,
  /return\s*\[\s*pathToTarget\s*,\s*(?:`Value found at index\s*\$\{\s*mid\s*\}`|['"]Value found at index\s*['"]\s*\.\s*concat\(\s*mid\s*\))\s*\]/
);
```

Ваша функція має повертати очікувані результати.

```js
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5], 3), [[3], "Value found at index 2"]
);
assert.deepEqual(
  binarySearch([1, 2, 3, 4, 5, 9], 4), [[3, 5, 4], "Value found at index 3"]
);
assert.deepEqual(
  binarySearch([1, 3, 5, 9, 14, 22], 10), [[], "Value not found"]
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
--fcc-editable-region--
      return pathToTarget;
--fcc-editable-region--
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```

# --solutions--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
    
    if (value === valueAtMiddle) {
      return [pathToTarget, `Value found at index ${mid}`];
    } else if (value > valueAtMiddle) {
      low = mid + 1;
    } else {
      high = mid - 1;
    }
  }
  return [[], "Value not found"];
}

console.log(binarySearch([1, 2, 3, 4, 5], 3));
console.log(binarySearch([1, 2, 3, 4, 5, 9], 4));
console.log(binarySearch([1, 3, 5, 9, 14, 22], 10));
```
