---
id: 6996c796385b6ea4e222d5e8
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Щоб почати цикл, потрібно знайти середній індекс поточного пошукового простору.

Ви можете знайти середній індекс, обчисливши середнє значення змінних `low` і `high`. Потрібно буде використати цілочисельне ділення після додавання двох значень, щоб отримати середнє, яке завжди буде цілим числом, округленим вниз, оскільки індекси мають бути цілими числами.

У вашому циклі `while` створіть змінну `mid`, яка дорівнює середньому значенню `low` і `high`.

# --hints--

У циклі while слід створити змінну `mid`, яка дорівнює `Math.floor((low + high) / 2)`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /(const|let|var)\s+mid\s*=\s*Math\.floor\(\s*\(\s*low\s*\+\s*high\s*\)\s*\/\s*2\s*\)/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
--fcc-editable-region--
      
--fcc-editable-region--
  }
}
```
