---
id: 6996c95e5bede03bdf2937a2
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Наступним кроком є створення умови, яка перевірятиме, чи знаходиться ціль `value` посередині.

Всередині циклу `while` створіть інструкцію `if`, яка перевіряє, чи дорівнює ціль `value` `valueAtMiddle`. Якщо так, поверніть змінну `pathToTarget`.

# --hints--

Ви маєте створити інструкцію `if` з умовою `value === valueAtMiddle`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)/
);
```

Ви маєте повернути змінну `pathToTarget` всередині інструкції `if`.

```js
assert.match(
  __helpers.removeJSComments(String(binarySearch)),
  /if\s*\(\s*value\s*===?\s*valueAtMiddle\s*\)\s*\{[\s\S]*?return\s+pathToTarget/
);
```

# --seed--

## --seed-contents--

```js
function binarySearch(searchList, value) {
  let pathToTarget = [];
  let low = 0;
  let high = searchList.length - 1;
  while (low <= high) {
    let mid = Math.floor((low + high) / 2);
    let valueAtMiddle = searchList[mid];
    pathToTarget.push(valueAtMiddle);
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
