---
id: 5895f70bf9fc0f352b528e64
title: Tumia nguvu za injini ya templeti
challengeType: 2
forumTopicId: 301567
dashedName: use-a-template-engines-powers
---

# --description--

Moja ya vipengele vikuu vya kutumia injini ya templeti ni uwezo wa kupitisha vigezo kutoka seva kwenda faili la templeti kabla ya kuonyesha kwa HTML.

Katika faili yako ya Pug, unaweza kutumia kigezo kwa kurejelea jina la kigezo kama `#{variable_name}` ndani ya mstari pamoja na maandishi mengine kwenye kipengele au kwa kutumia ishara ya usawa kwenye kipengele bila nafasi kama `p=variable_name` ambayo inaweka thamani ya kigezo kama maandishi ya kipengele cha p.

Pug inahusu kabisa kutumia nafasi tupu na tabo kuonyesha vipengele vilivyopangwa ndani na kupunguza kiasi cha msimbo unaohitajika kutengeneza tovuti nzuri.

Chukua msimbo ufuatao wa Pug kama mfano:

```pug
head
  script(type='text/javascript').
    if (foo) bar(1 + 5);
body
  if youAreUsingPug
      p You are amazing
    else
      p Get on it!
```

Msimbo huo hapo juu hutengeneza HTML ifuatayo:

```html
<head>
  <script type="text/javascript">
    if (foo) bar(1 + 5);
  </script>
</head>
<body>
  <p>You are amazing</p>
</body>
```

Faili lako la `index.pug` lililomo katika mradi wako, linatumia vigezo `title` na `message`.

Pitisha hivyo kutoka seva yako kwenda faili la Pug kwa kuongeza kitu kama hoja ya pili kwenye wito wako wa `res.render` lenye vigezo na thamani zao. Mpe `title` thamani ya `Hello` na `message` thamani ya `Please log in`.

Itakuwa kama ifuatavyo:

```javascript
res.render('index', { title: 'Hello', message: 'Please log in' });
```

Sasa onyesha upya ukurasa wako, na unapaswa kuona thamani hizo zikiwekwa kwenye muonekano wako mahali sahihi kama ilivyoainishwa katika faili lako la `index.pug`!

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza kuangalia <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#use-a-template-engines-power-2" target="_blank" rel="noopener noreferrer nofollow">mradi uliofikia hatua hii</a>.

# --hints--

Pug inapaswa kuonyesha vigezo kwa usahihi.

```js
  const url = new URL("/", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /pug-variable("|')>Please log in/gi,
    'Your projects home page should now be rendered by pug with the projects .pug file unaltered'
  );
```
