---
id: 5895f70ef9fc0f352b528e6b
title: Jinsi ya Kuandaa Profaili
challengeType: 2
forumTopicId: 301554
dashedName: how-to-put-a-profile-together
---

# --description--

Sasa unaweza kuhakikisha mtumizi anayeupata upatikanaji wa `/profile` amethibitishwa, unaweza kutumia taarifa zilizomo katika `req.user` kwenye ukurasa wako.

Pitia kitu chenye sifa ya `username` na thamani ya `req.user.username` kama hoja ya pili kwa njia ya `render` ya mtazamo wa profaili.

Kisha, nenda kwenye mtazamo wako wa `profile.pug`, na ongeza mstari ufuatao chini ya kipengele cha `h1` kilichopo, na kwa kiwango sawa cha kuingiza:

```pug
h2.center#welcome Welcome, #{username}!
```

Hii inaunda kipengele cha `h2` chenye darasa la vitu `center` na kitambulisho `welcome` chenye maandishi ya `Welcome, ` ikifuatiwa na jina la mtumizi.

Pia, katika `profile.pug`, ongeza kiungo kinachorejelea njia ya data ya `/logout`, ambayo itakuwa mwenyeji wa mantiki ya kumfanya mtumizi asithibitishwe:

```pug
a(href='/logout') Logout
```

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#how-to-put-a-profile-together-9" target="_blank" rel="noopener noreferrer nofollow">kagua mradi uliofanyika hadi hatua hii</a>.

# --hints--

Unapaswa kuongeza kwa usahihi kigezo cha onyesho la Pug katika `/profile`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /username:( |)req.user.username/,
    'You should be passing the variable username with req.user.username into the render function of the profile page'
  );
```
