---
id: 5895f70cf9fc0f352b528e66
title: Uwekaji mfululizo wa kitu cha mtumizi
challengeType: 2
forumTopicId: 301563
dashedName: serialization-of-a-user-object
---

# --description--

Uwekaji mfululizo na kuondoa mfululizo ni dhana muhimu kuhusu uthibitisho wa utambulisho. Kuweka kitu mfululizo kunamaanisha kubadilisha yaliyomo kuwa *funguo* ndogo ambayo baadaye inaweza kuondolewa mfululizo kuwa kitu halisi. Hii ndiyo inayotuwezesha kujua ni nani aliyewasiliana na seva bila ya kutuma data za uthibitisho wa utambulisho, kama jina la mtumizi na nenosiri, kila ombi la ukurasa mpya.

Ili kuandaa hili ipasavyo, unahitaji kuwa na kitendakazi cha kuweka mfululizo na kitendakazi cha kuondoa mfululizo. Katika Passport, hizi zinaweza kuundwa kwa:

```javascript
passport.serializeUser(cb);
passport.deserializeUser(cb);
```

Kitendakazi cha mwitiko kinachopitishwa kwa `serializeUser` kinaitwa na hoja mbili: kitu kamili cha mtumizi, na wito wa kurudisha unaotumiwa na passport.

Wito wa kurudisha unatarajia hoja mbili: Hitilafu, ikiwa ipo, na funguo ya kipekee ya kumtambua mtumizi ambayo inapaswa kurudishwa katika wito wa kurudisha. Utatumia `_id` ya mtumizi katika kitu hicho. Hii inahakikishwa kuwa ni ya kipekee, kwani inazalishwa na MongoDB.

Vivyo hivyo, `deserializeUser` inaitwa na hoja mbili: funguo ya kipekee, na kitendakazi cha mwitiko.

Kitendakazi hiki cha mwitiko kinatarajia hoja mbili: Hitilafu, ikiwa ipo, na kitu kamili cha mtumizi. Ili kupata kitu kamili cha mtumizi, fanya hoja ya utafutaji kwa Mongo `_id`, kama inavyoonyeshwa hapa chini:

```javascript
passport.serializeUser((user, done) => {
  done(null, user._id);
});

passport.deserializeUser((id, done) => {
  myDataBase.findOne({ _id: new ObjectID(id) }, (err, doc) => {
    done(null, null);
  });
});
```

Ongeza vitendakazi viwili hapo juu kwenye seva yako. Darasa la `ObjectID` linatoka kwenye kifurushi cha `mongodb`. `mongodb@~3.6.0` tayari imeongezwa kama utegemezi. Tangaza darasa hili kwa:

```javascript
const { ObjectID } = require('mongodb');
```

`deserializeUser` itatupa hitilafu hadi utakapoweka muunganisho wa hifadhidata. Kwa sasa, toa maoni kwenye wito wa `myDatabase.findOne`, na piga tu wito wa `done(null, null)` katika kitendakazi cha mwitiko cha `deserializeUser`.

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na hitilafu, unaweza <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#serialization-of-a-user-object-4" target="_blank" rel="noopener noreferrer nofollow">kagua mradi uliofanyika hadi hatua hii</a>.

# --hints--

Unapaswa kuweka kitu cha mtumizi mfululizo ipasavyo.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.serializeUser/gi,
    'You should have created your passport.serializeUser function'
  );
  assert.match(
    data,
    /null,\s*user._id/gi,
    'There should be a callback in your serializeUser with (null, user._id)'
  );
```

Unapaswa kuondoa mfululizo wa kitu cha mtumizi ipasavyo.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /passport.deserializeUser/gi,
    'You should have created your passport.deserializeUser function'
  );
  assert.match(
    data,
    /null,\s*null/gi,
    'There should be a callback in your deserializeUser with (null, null) for now'
  );
```

MongoDB inapaswa kuwa utegemezi.

```js
  const url = new URL("/_api/package.json", code);
  const res = await fetch(url);
  const packJson = await res.json();
  assert.property(
    packJson.dependencies,
    'mongodb',
    'Your project should list "mongodb" as a dependency'
  );
```

MongoDB inapaswa kuhitajika ipasavyo ikiwa ni pamoja na ObjectId.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /require.*("|')mongodb\1/gi,
    'You should have required mongodb'
  );
  assert.match(
    data,
    /new ObjectID.*id/gi,
    'Even though the block is commented out, you should use new ObjectID(id) for when we add the database'
  );
```
