---
id: 58965611f9fc0f352b528e6c
title: Kuingia nje kwa mtumizi
challengeType: 2
forumTopicId: 301560
dashedName: logging-a-user-out
---

# --description--

Kuunda mantiki ya kuingia nje ni rahisi. Njia ya data inapaswa tu kuondoa uthibitisho wa mtumizi, na kuelekeza upya kwenye ukurasa wa nyumbani badala ya kuonyesha mtazamo wowote.

Katika passport, kuondoa uthibitisho wa mtumizi ni rahisi kama tu kufanya wito wa `req.logout()` kabla ya kuelekeza upya. Ongeza njia ya data ya `/logout` kufanya hivyo:

```js
app.route('/logout')
  .get((req, res) => {
    req.logout();
    res.redirect('/');
});
```

Huenda umeona kuwa haujadhibiti kurasa zilizopotea (404). Njia ya kawaida ya kushughulikia hili katika nodi ni kwa kutumia programu ya kati ifuatayo. Endelea na ongeza hii baada ya njia zako zote za data:

```js
app.use((req, res, next) => {
  res.status(404)
    .type('text')
    .send('Not Found');
});
```

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#logging-a-user-out-10" target="_blank" rel="noopener noreferrer nofollow">kagua mradi uliofikia hatua hii</a>.

# --hints--

`req.logout()` inapaswa kuitwa katika njia yako ya data ya `/logout`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /req.logout/gi,
    'You should be calling req.logout() in your /logout route'
  );
```

`/logout` inapaswa kuelekeza upya kwenye ukurasa wa nyumbani.

```js
  const url = new URL("/logout", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /Home page/gi,
    'When a user logs out they should be redirected to the homepage'
  );
```
