---
id: 589a8eb3f9fc0f352b528e72
title: Utekelezaji wa uthibitisho wa utambulisho wa kijamii III
challengeType: 2
forumTopicId: 301558
dashedName: implementation-of-social-authentication-iii
---

# --description--

Sehemu ya mwisho ya mkakati ni kushughulikia wasifu unaorudishwa kutoka GitHub. Tunahitaji kupakia kitu cha mtumiaji katika hifadhidata ikiwa kipo, au kuunda kipya ikiwa hakipo, na kujaza sehemu kutoka kwenye wasifu, kisha kurudisha kitu cha mtumiaji. GitHub hutupatia *id* ya kipekee ndani ya kila wasifu ambayo tunaweza kuitumia kutafuta ili kusanifu mtumiaji (ambayo tayari imetekelezwa). Hapa chini ni mfano wa utekelezaji unaweza kutumia katika mradi wako—unaenda ndani ya kitendakazi ambacho ni hoja ya pili kwa mkakati mpya, chini kabisa ya mahali ambapo `console.log(profile);` iko sasa:

```js
myDataBase.findOneAndUpdate(
  { id: profile.id },
  {
    $setOnInsert: {
      id: profile.id,
      username: profile.username,
      name: profile.displayName || 'John Doe',
      photo: profile.photos[0].value || '',
      email: Array.isArray(profile.emails)
        ? profile.emails[0].value
        : 'No public email',
      created_on: new Date(),
      provider: profile.provider || ''
    },
    $set: {
      last_login: new Date()
    },
    $inc: {
      login_count: 1
    }
  },
  { upsert: true, new: true },
  (err, doc) => {
    return cb(null, doc.value);
  }
);
```

`findOneAndUpdate` inakuwezesha kutafuta kitu na kukiboresha. Ikiwa kitu hakipo, kitaingizwa na kutolewa kwa kitendakazi cha mwitiko. Katika mfano huu, kila mara tunaweka `last_login`, kuongezea `login_count` kwa `1`, na kujaza sehemu nyingi tu wakati kitu kipya (mtumiaji mpya) kinaingizwa. Angalia matumizi ya thamani za msingi. Wakati mwingine wasifu unaorudishwa hautakuwa na taarifa zote zimejazwa au mtumiaji ataweka kuwa faragha. Katika hali hii, unashughulikia ili kuzuia kosa.

Sasa unapaswa kuweza kuingia kwenye programu yako. Jaribu!

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#implementation-of-social-authentication-iii-5" target="_blank" rel="noopener noreferrer nofollow">kuangalia mradi uliofanyika hadi hatua hii</a>.

# --hints--

Usanidi wa mkakati wa GitHub unapaswa kuwa umekamilika.

```js
  const url = new URL("/_api/auth.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /GitHubStrategy[^]*myDataBase/gi,
    'Strategy should use now use the database to search for the user'
  );
  assert.match(
    data,
    /GitHubStrategy[^]*cb/gi,
    'Strategy should return the callback function "cb"'
  );
```
