---
id: 589fc831f9fc0f352b528e76
title: Shughulikia kutengana kwa muunganisho
challengeType: 2
forumTopicId: 301552
dashedName: handle-a-disconnect
---

# --description--

Unaweza kugundua kwamba hadi sasa umekuwa ukiongeza tu idadi ya watumizi. Kushughulikia mtumizi kutengana ni rahisi kama kushughulikia muunganisho wa awali, isipokuwa lazima usikilize tukio hilo kwenye kila nodi badala ya kwenye seva nzima.

Ili kufanya hivi, ongeza msikilizaji mwingine ndani ya msikilizaji uliopo wa `'connect'` unaosikiliza tukio la `'disconnect'` kwenye nodi bila data yoyote kupitishwa. Unaweza kujaribu utendaji huu kwa kurekodi tu kwenye konsoli kwamba mtumizi amejitenga.

```js
socket.on('disconnect', () => {
  /*anything you want to do on disconnect*/
});
```

Ili kuhakikisha wateja wanapata idadi iliyosasishwa ya watumizi wa sasa kila wakati, unapaswa kupunguza `currentUsers` kwa 1 wakati kutengana kunapotokea kisha toa tukio la `'user count'` lenye idadi iliyosasishwa.

**Kumbuka:** Kama ilivyo kwa `'disconnect'`, matukio mengine yote ambayo nodi inaweza kuyatoa kwa seva yanapaswa kushughulikiwa ndani ya msikilizaji wa muunganisho ambapo 'socket' imefafanuliwa.

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135#handle-a-disconnect-8" target="_blank" rel="noopener noreferrer nofollow">kuangalia mradi uliofanyika hadi hatua hii</a>.

# --hints--

Seva inapaswa kushughulikia tukio la disconnect kutoka kwenye nodi.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(data, /socket.on.*('|")disconnect('|")/s, '');
```

Mteja wako anapaswa kusikiliza tukio la `'user count'`.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user count('|")/s,
    'Your client should be connection to server with the connection defined as socket'
  );
```
