---
id: 589fc832f9fc0f352b528e78
title: Tangaza watumizi wapya
challengeType: 2
forumTopicId: 301546
dashedName: announce-new-users
---

# --description--

Sehemu nyingi za mazungumzo zina uwezo wa kutangaza wakati mtumizi anapo ungana au kutenganishwa na kisha kuonyesha hilo kwa watumizi wote waliounganishwa kwenye mazungumzo. Kwa kuwa tayari unatoa tukio wakati wa kuungana na kutenganishwa, utahitaji tu kubadilisha tukio hili ili kuunga mkono kipengele kama hicho. Njia yenye mantiki zaidi ni kutuma vipande 3 vya data pamoja na tukio: jina la mtumizi aliyekuwa anajiunga/akatenganishwa, idadi ya watumizi waliopo sasa, na kama jina hilo la mtumizi lilijiunga au kutenganishwa.

Badilisha jina la tukio kuwa `'user'`, na pita kitu chenye sehemu `username`, `currentUsers`, na `connected` (kitakuwa `true` ikiwa ni kuungana, au `false` kwa kutenganishwa kwa mtumizi aliyepelekwa). Hakikisha kubadilisha matukio yote ya `'user count'` na weka tukio la kutenganishwa litume `false` kwa sehemu ya `connected` badala ya `true` kama tukio lililotolewa wakati wa kuungana.

```js
io.emit('user', {
  username: socket.request.user.username,
  currentUsers,
  connected: true
});
```

Sasa mteja wako atakuwa na taarifa zote muhimu za kuonyesha kwa usahihi idadi ya watumizi waliopo sasa na kutangaza wakati mtumizi anajiunga au kutenganishwa! Ili kushughulikia tukio hili upande wa mteja tunapaswa kusikiliza `'user'`, kisha sasisha idadi ya watumizi waliopo kwa kutumia jQuery kubadilisha maandishi ya `#num-users` kuwa `'{NUMBER} users online'`, pamoja na kuongeza `<li>` kwenye orodha isiyopangwa yenye kitambulisho `messages` na `'{NAME} has {joined/left} the chat.'`.

Utekelezaji wa hili unaweza kuonekana kama ifuatavyo:

```js
socket.on('user', data => {
  $('#num-users').text(data.currentUsers + ' users online');
  let message =
    data.username +
    (data.connected ? ' has joined the chat.' : ' has left the chat.');
  $('#messages').append($('<li>').html('<b>' + message + '</b>'));
});
```

Wasilisha ukurasa wako unapoona umefanya vizuri. Ikiwa unakutana na makosa, unaweza kuangalia <a href="https://forum.freecodecamp.org/t/advanced-node-and-express/567135/3#announce-new-users-10" target="_blank" rel="noopener noreferrer nofollow">mradi uliofanyika hadi hatua hii</a>.

# --hints--

Tukio `'user'` linapaswa kutolewa na `username`, `currentUsers`, na `connected`.

```js
  const url = new URL("/_api/server.js", code);
  const res = await fetch(url);
  const data = await res.text();
  // Regex is lenient to match both `username` and `name` as the key on purpose.
  assert.match(
    data,
    /io.emit.*('|")user\1.*name.*currentUsers.*connected/s,
    'You should have an event emitted named user sending name, currentUsers, and connected'
  );
```

Mteja anapaswa kushughulikia na kuonyesha data mpya kutoka kwa tukio `'user'` kwa usahihi.

```js
  const url = new URL("/public/client.js", code);
  const res = await fetch(url);
  const data = await res.text();
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*num-users/s,
    'You should change the text of "#num-users" within on your client within the "user" event listener to show the current users connected'
  );
  assert.match(
    data,
    /socket.on.*('|")user\1[^]*messages.*li/s,
    'You should append a list item to "#messages" on your client within the "user" event listener to announce a user came or went'
  );
```
