---
id: 613e275749ebd008e74bb62e
title: Paso 8
challengeType: 0
dashedName: step-8
---

# --description--

Una propiedad útil de un _SVG_ (gráfico vectorial escalable) es que contiene un atributo `path`, el cual permite escalar la imagen sin afectar la resolución de la imagen resultante.

Actualmente, `img` asume su tamaño predeterminado, que es demasiado grande. CSS tiene una función `max`, que devuelve el elemento más grande de un conjunto de valores separados por comas. Por ejemplo:

```css
img {
  width: max(250px, 25vw);
}
```

En el ejemplo anterior, el ancho de la imagen será 250px si el ancho de la vista es inferior a 1000 píxeles. Si el ancho de la vista es mayor que 1000 píxeles, el ancho de la imagen será 25vw. Esto ocurre porque 25vw es igual que 25% del ancho de la vista.

Escala la imagen usando su `id` como selector y estableciendo el valor de `width` al máximo de `10rem` o `18vw`.

# --hints--

Debes usar el selector `#logo` para apuntar al elemento `img`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('#logo'));
```

Debes dar `img` un `width` de `max(10rem, 18vw)`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#logo')?.width, 'max(10rem, 18vw)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav></nav>
    </header>
    <main></main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

--fcc-editable-region--

--fcc-editable-region--

```
