---
id: a5229172f011153519423690
title: Construye un Calculador de Suma de Fibonacci Impar
challengeType: 26
dashedName: build-an-odd-fibonacci-sum-calculator
---

# --description--

En este laboratorio construirás un calculador de suma de Fibonacci impar que calcula la suma de todos los números de Fibonacci impares que son menores o iguales a un entero positivo dado.

**Objetivo**: Cumple las historias de usuario a continuación y logra que todas las pruebas pasen para completar el laboratorio.

**Historias de Usuario**

1. Debes tener una función `sumFibs` que acepte un número como argumento.
1. La función `sumFibs` debe devolver la suma de todos los números impares de Fibonacci que sean menores o iguales al número dado.
1. La secuencia de Fibonacci comienza con `0` y `1`, y cada número siguiente es la suma de los dos anteriores.
1. Solo los números impares de Fibonacci deben sumarse a la suma.

# --hints--

Deberías tener una función `sumFibs`.

```js
assert.isFunction(sumFibs);
```

`sumFibs(1)` debe devolver un número.

```js
assert.isNumber(sumFibs(1));
```

`sumFibs(1000)` debería devolver `1785`.

```js
assert.strictEqual(sumFibs(1000), 1785);
```

`sumFibs(4000000)` debería devolver `4613732`.

```js
assert.strictEqual(sumFibs(4000000), 4613732);
```

`sumFibs(4)` debería devolver `5`.

```js
assert.strictEqual(sumFibs(4), 5);
```

`sumFibs(75024)` debería devolver `60696`.

```js
assert.strictEqual(sumFibs(75024), 60696);
```

`sumFibs(75025)` debería devolver `135721`.

```js
assert.strictEqual(sumFibs(75025), 135721);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumFibs(num) {
  let a = 0, b = 1;
  let sum = 0;

  while (b <= num) {
    if (b % 2 !== 0) {
      sum += b;
    }
    [a, b] = [b, a + b];
  }

  return sum;
}
```
