---
id: 68e6a99a0fd7386fb66e156f
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Katika warsha hii, utazoea kupamba kisanduku cha kuchagua kwa kujenga paneli ya mipangilio ya mchezo.

Muundo wa msingi wote wa HTML (`DOCTYPE`, `html`, `head` na `body`) umeandaliwa kwa ajili yako.

Ndani ya `body`, tengeneza kipengele cha `div` chenye kipengele cha `h1` kilichopangwa ndani. Mpe kipengele cha `h1` maandishi ya `Game Settings`

# --hints--

Unapaswa kuwa na kipengele cha `div`.

```js
assert.lengthOf(document.querySelectorAll("div"), 1);
```

Unapaswa kuwa na kipengele cha `h1`.

```js
assert.lengthOf(document.querySelectorAll("h1"), 1);
```

Kipengele chako cha `h1` kinapaswa kuwa ndani ya kipengele chako cha `div`.

```js
assert.exists(document.querySelector("div")?.querySelector("h1"));
```

Kipengele chako cha `h1` kinapaswa kuwa na maandishi ya `Game Settings`.

```js
assert.match(document.querySelector("h1")?.textContent, /Game Settings/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
