---
id: 68e9a7bd60c971278d2d4ffc
title: Hatua 13
challengeType: 0
dashedName: step-13
---

# --description--

Katika funzo lililopita, ulijifunza kuhusu madarasa ya bandia na vipengele vya bandia katika CSS. Sasa utaweka maarifa hayo kwa kuunda kichaguzi cha aina kilichojumuishwa na kichaguzi cha darasa la bandia.

Kwanza, anza na muundo wa kichaguzi cha aina kwa kukiweka kama `input[type="checkbox"]` na kuongeza `:checked` mwishoni mwao.

Kisha, kipa `background-color` yenye thamani ya `#f1be32`.

Na hatimaye, weka `border-color` kuwa na thamani ya `#e2a60d`.

# --hints--

Unapaswa kuwa na kichaguzi kilichojumuisha aina na darasa la bandia cha `input[type="checkbox"]:checked`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked'));
```

Unapaswa kuongeza sifa ya `background-color` kwa kichaguzi chako kilichojumuisha aina na darasa la bandia cha `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor);
```

Unapaswa kuwa na sifa ya `background-color` yenye thamani ya `#f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.backgroundColor, "rgb(241, 190, 50)");
```

Unapaswa kuongeza sifa ya `border-color` kwa kichaguzi chako kilichojumuisha aina na darasa la bandia cha `input[type="checkbox"]:checked`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor);
```

Unapaswa kuwa na sifa ya `border-color` yenye thamani ya `#e2a60d`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked')?.borderColor, "rgb(226, 166, 13)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

--fcc-editable-region--

--fcc-editable-region--

```
