---
id: 68e9a834dfef752dc35f6f90
title: Hatua 14
challengeType: 0
dashedName: step-14
---

# --description--

Sasa baada ya kisanduku chako cha kuchagua kubadilisha rangi wakati mtumizi anabofya, utaweka mtindo fulani ili kuonyesha kwa uwazi kidogo kwamba kisanduku cha kuchagua kimechaguliwa.

Darasa la bandia la CSS linaweza kuunganishwa. Unaweza kuongeza darasa la bandia au kipengele cha bandia kifuatacho ili kulenga vipengele vinavyokidhi masharti mengi.

Tengeneza kichaguzi kipya kwa `input[type="checkbox"]`. Ongeza darasa la bandia la `:checked` ikifuatiwa na kipengele cha bandia cha `::after`. Hii inakuwezesha kuongeza kiashiria cha kuona wakati kisanduku cha kuchagua kimechaguliwa.

Kisha, ongeza `content` yenye thamani ya `"✓"`.

# --hints--

Unapaswa kuwa na kichaguzi cha aina kilichounganishwa na kichaguzi cha darasa la bandia na kichaguzi cha kipengele cha bandia cha `input[type="checkbox"]:checked::after`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after'));
```

Unapaswa kuongeza sifa ya `content` kwa kichaguzi chako cha aina kilichounganishwa na kichaguzi cha darasa la bandia na kichaguzi cha kipengele cha bandia cha `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content);
```

Unapaswa kuwa na sifa ya `content` yenye thamani ya `"✓"`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.content, '"✓"');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

--fcc-editable-region--

--fcc-editable-region--

```
