---
id: 68e9a3194160f8f5620ddf21
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Mara moja baada ya kipengele chako cha `h1`, tengeneza vipengele vinne vya `label`.

Ndani ya kila kipengele cha `label`, tengeneza kipengele cha `input` chenye sifa ya kipengele ya `type` iliyowekwa kuwa `"checkbox"`.

Baada ya kila kipengele cha `input`, toa maandishi yafuatayo kwa vipengele vyako vya `label`, kwa mpangilio: `Sound Effects`, `Background Music`, `Hard Mode`, na `Haptic Feedback`.

# --hints--

Unapaswa kuwa na vipengele vinne vya `label`.

```js
assert.lengthOf(document.querySelectorAll("label"), 4);
```

Kipengele chako cha kwanza cha `label` kinapaswa kuwa na kipengele cha `input`.

```js
assert.exists(document.querySelectorAll("label")[0]?.querySelector("input"));
```

Kipengele chako cha kwanza cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[0]?.querySelector("input")?.type, "checkbox");
```

Kipengele chako cha kwanza cha `label` kinapaswa kuwa na maandishi ya `Sound Effects`.

```js
assert.match(document.querySelectorAll("label")[0]?.textContent, /Sound Effects/i);
```

Maandishi yako ya `Sound Effects` yanapaswa kuja baada ya kipengele cha `input`.

```js
assert.match(document.querySelectorAll("label")[0]?.innerHTML, /\<input[^>]*>\s*Sound Effects/i)
```

Kipengele chako cha pili cha `label` kinapaswa kuwa na kipengele cha `input`.

```js
assert.exists(document.querySelectorAll("label")[1]?.querySelector("input"));
```

Kipengele chako cha pili cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[1]?.querySelector("input")?.type, "checkbox");
```

Kipengele chako cha pili cha `label` kinapaswa kuwa na maandishi ya `Background Music`.

```js
assert.match(document.querySelectorAll("label")[1]?.textContent, /Background Music/i)
```

Maandishi yako ya `Background Music` yanapaswa kuja baada ya kipengele cha `input`.

```js
assert.match(document.querySelectorAll("label")[1]?.innerHTML, /\<input[^>]*>\s*Background Music/i)
```

Kipengele chako cha tatu cha `label` kinapaswa kuwa na kipengele cha `input`.

```js
assert.exists(document.querySelectorAll("label")[2]?.querySelector("input"));
```

Kipengele chako cha tatu cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[2]?.querySelector("input")?.type, "checkbox");
```

Kipengele chako cha tatu cha `label` kinapaswa kuwa na maandishi ya `Hard Mode`.

```js
assert.match(document.querySelectorAll("label")[2]?.textContent, /Hard Mode/i)
```

Maandishi yako ya `Hard Mode` yanapaswa kuja baada ya kipengele cha `input`.

```js
assert.match(document.querySelectorAll("label")[2]?.innerHTML, /\<input[^>]*>\s*Hard Mode/i)
```

Kipengele chako cha nne cha `label` kinapaswa kuwa na kipengele cha `input`.

```js
assert.exists(document.querySelectorAll("label")[3]?.querySelector("input"));
```

Kipengele chako cha nne cha `input` kinapaswa kuwa na sifa ya kipengele ya `type` iliyowekwa kuwa `checkbox`.

```js
assert.strictEqual(document.querySelectorAll("label")[3]?.querySelector("input")?.type, "checkbox");
```

Kipengele chako cha nne cha `label` kinapaswa kuwa na maandishi ya `Haptic Feedback`.

```js
assert.match(document.querySelectorAll("label")[3]?.textContent, /Haptic Feedback/i)
```

Maandishi yako ya `Haptic Feedback` yanapaswa kuja baada ya kipengele cha `input`.

```js
assert.match(document.querySelectorAll("label")[3]?.innerHTML, /\<input[^>]*>\s*Haptic Feedback/i)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
--fcc-editable-region--
      
--fcc-editable-region--
    </div>
  </body>
</html>
```
