---
id: 68e9a6bf1a74721a4e43e217
title: Hatua 10
challengeType: 0
dashedName: step-10
---

# --description--

Sasa utaondoa kisanduku cha kuchagua cha kawaida kinachotumika na vivinjari. Ndani ya kichaguzi chako cha `input[type='checkbox']`, ongeza `appearance` yenye thamani ya `none`.

Kuweka sifa ya `appearance` kuwa `none` kutafuta muonekano ambao kivinjari hutumia kwa visanduku vya kuchagua, na kukuwezesha kuonyesha mtindo unaotaka.

Baada ya kufanya hivyo, kwa kuwa kisanduku cha kuchagua hakitaonekana tena, weka `border` yenye unene wa `2px`, mtindo wa `solid`, na nambari ya hex ya `#f1be32`.

# --hints--

Unapaswa kuongeza sifa ya `appearance` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance);
```

Unapaswa kuwa na sifa ya `appearance` yenye thamani ya `none`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.appearance, "none");
```

Unapaswa kuongeza sifa ya `border` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border);
```

Unapaswa kuwa na sifa ya `border` yenye thamani ya `2px solid #f1be32`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.border, "2px solid rgb(241, 190, 50)");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
