---
id: 68e9a6201e2efa142d62bfe1
title: Hatua 9
challengeType: 0
dashedName: step-9
---

# --description--

Hatua inayofuata, utaongeza ukubwa wa visanduku vya kuchagua kwa ajili ya uonekano bora.

Anza kwa kuandaa kichaguzi cha `input`, lakini ukilenga hasa `[type="checkbox"]` yako.

Ndani ya kichaguzi chako, weka `width` na `height` kuwa `20px`. Hii inafanya ziwe kubwa kuliko zilivyokuwa awali.

Na ili kuendana na mpangilio wako wa `cursor` uliowekwa kwenye lebo, ongeza `cursor` na weka thamani ya `pointer` kwake. Baada ya hapo, unapobeba kidole juu ya visanduku vya kuchagua, itaonyesha kidokezo cha pointer.

# --hints--

Unapaswa kuwa na kichaguzi katika CSS yako cha `input[type="checkbox"]`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]'));
```

Unapaswa kuongeza sifa ya `width` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width);
```

Unapaswa kuwa na sifa ya `width` yenye thamani ya `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.width, "20px");
```

Unapaswa kuongeza sifa ya `height` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height);
```

Unapaswa kuwa na sifa ya `height` yenye thamani ya `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.height, "20px");
```

Unapaswa kuongeza sifa ya `cursor` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor);
```

Unapaswa kuwa na sifa ya `cursor` yenye thamani ya `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

--fcc-editable-region--

--fcc-editable-region--
```
