---
id: 68ee86e3ee07b01626a392cc
title: Hatua 6
challengeType: 0
dashedName: step-6
---

# --description--

Ndani ya kichaguzi chako cha `.settings-card` weka sifa ya `margin` kuwa `auto`.

Kuweka sifa ya `margin` kuwa `auto` hujirekebisha kiotomatiki ukingo wa kipengee ili kusambaza nafasi iliyobaki sawasawa ndani ya kontena lake, mara nyingi hutumika kuweka vitu vya darasa la block katikati kwa usawa wa mlalo.

Na mwisho, weka sifa ya `text-align` yenye thamani ya `left`. Hii itaweka maudhui ya ndani ya mstari, kama maandishi, upande wa kushoto wa kipengee kinachoshikilia.

# --hints--

Unapaswa kuongeza sifa ya `margin` kwa kichaguzi cha darasa la `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin);
```

Unapaswa kuwa na sifa ya `margin` yenye thamani ya `auto`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.margin, "auto");
```

Unapaswa kuongeza sifa ya `text-align` kwa kichaguzi cha darasa la `settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign);
```

Unapaswa kuwa na sifa ya `text-align` yenye thamani ya `left`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.textAlign, "left");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
