---
id: 68f15cd397503429f0345262
title: Hatua 15
challengeType: 0
dashedName: step-15
---

# --description--

Katika tamko la kichaguzi cha `input[type="checkbox"]:checked::after`, weka `display` kuwa `block`. Kuweka sifa yako ya `display` kama `block` hufanya kipengele kuwa kipengele cha darasa la vitu cha ngazi ya blok, ikimaanisha kwamba kinachukua upana mzima wa kontena lake na kuanza kwenye mstari mpya, kukuwezesha kudhibiti upana, urefu, na nafasi kwa urahisi zaidi.

Ifuatayo, weka `text-align` kuwa `center`. Hii italeta maudhui ya ndani ya mstari (kama maandishi, au vipengele vya mstari mmoja) katikati kwa usawa ndani ya blok. Blok bado itachukua upana mzima wa kontena lake lakini kila kitu ndani yake kitawekwa katikati.

# --hints--

Unapaswa kuongeza sifa ya `display` kwa kichaguzi chako cha aina kilichochanganywa na kichaguzi cha darasa la bandia na kichaguzi cha kipengele cha bandia cha `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display);
```

Unapaswa kuwa na sifa ya `display` yenye thamani ya `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.display, "block");
```

Unapaswa kuongeza sifa ya `text-align` kwa kichaguzi chako cha aina kilichochanganywa na kichaguzi cha darasa la bandia na kichaguzi cha kipengele cha bandia cha `input[type="checkbox"]:checked::after`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign);
```

Unapaswa kuwa na sifa ya `text-align` yenye thamani ya `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]:checked::after')?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
  vertical-align: middle;
  margin-right: 15px;
}

input[type='checkbox']:checked {
  background-color: #f1be32;
  border-color: #e2a60d;
}

input[type="checkbox"]:checked::after {
  content: "✓";
--fcc-editable-region--

--fcc-editable-region--
}
```
