---
id: 68f1560aff1214103ab6b59b
title: Hatua 3
challengeType: 0
dashedName: step-3
---

# --description--

Sasa utaanza kupamba ukurasa kwa mtindo wa CSS. Anza kwa kuunda kichaguzi cha `body`.

Weka sifa ya `body` kuwa na sifa ya `height` yenye thamani ya `100vh` na sifa ya `background-color` yenye thamani ya `#f0f0f0`.

`height` ya `100vh` hufanya `body` kuchukua urefu mzima wa eneo la kuona la kivinjari, wakati rangi ya asili ya kijivu nyepesi hutoa msingi laini kwa ukurasa.

Mwishowe, weka sifa ya `text-align` yenye thamani ya `center`. Hii italinganisha katikati maudhui yote yaliyomo ndani ya mstari ndani ya ukurasa isipokuwa kipengele cha mtoto kitaghairi na ulinganifu wake mwenyewe.

# --hints--

Unapaswa kutumia kichaguzi cha aina cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("body"));
```

Unapaswa kuongeza sifa ya `height` ndani ya kichaguzi cha aina cha `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.height);
```

Unapaswa kuwa na sifa ya `height` yenye thamani ya `100vh`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.height, "100vh");
```

Unapaswa kuongeza sifa ya `background-color` ndani ya kichaguzi cha aina cha `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor);
```

Unapaswa kuwa na sifa ya `background-color` yenye thamani ya `#f0f0f0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("body")?.backgroundColor, "rgb(240, 240, 240)");
```

Unapaswa kuongeza sifa ya `text-align` ndani ya kichaguzi cha aina cha `body`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("body")?.textAlign);
```

Unapaswa kuwa na sifa ya `text-align` yenye thamani ya `center`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("body")?.textAlign, "center");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div>
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
