---
id: 68f15ef0446464325ce1aa45
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--

Unahitaji nafasi kati ya visanduku vya kuchagua na lebo. Anza kwa kutumia kichaguzi cha `label` katika CSS yako na weka sifa ya `display` kuwa `block`.

Kisha, weka sifa ya `margin` kuwa `8px auto` kuongeza nafasi ya wima kati ya vipengele. Hii itaunda nafasi thabiti juu na chini ya kila kipengele, kusaidia kutenganisha maudhui na kuboresha usomaji.

Mwishowe, weka sifa ya `cursor` kuwa `pointer` kwenye vipengele vya `label`. Hii itabadilisha mshale wa kidokezo kuwa ikoni ya mkono wakati mtumizi anapobeba kidole juu ya `label`, ikionyesha kuwa kipengele kinaweza kubofywa na kuboresha uzoefu wa mtumizi kwa ujumla.

# --hints--

Unapaswa kuwa na kichaguzi cha aina katika CSS yako kwa `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle("label"));
```

Unapaswa kuongeza sifa ya `display` kwa kichaguzi chako cha `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.display);
```

Unapaswa kuwa na sifa ya `display` yenye thamani ya `block`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.display, "block");
```

Unapaswa kuongeza sifa ya `margin` kwa kichaguzi chako cha `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.margin);
```

Unapaswa kuwa na sifa ya `margin` yenye thamani ya `8px auto`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle("label")?.margin, "8px auto");
```

Unapaswa kuongeza sifa ya `cursor` kwa kichaguzi chako cha `label`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle("label")?.cursor);
```

Unapaswa kuwa na sifa ya `cursor` yenye thamani ya `pointer`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle("label")?.cursor, "pointer");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}
--fcc-editable-region--

--fcc-editable-region--
```
