---
id: 68f29753fd5524530c6e65c8
title: Hatua 12
challengeType: 0
dashedName: step-12
---

# --description--

Ifuatayo, weka sifa ya `vertical-align` yenye thamani ya `middle`.

Sifa ya `vertical-align` inasimamia jinsi vipengele vya `inline` au `inline-block` vinavyopangwa wima na maandishi yanayozunguka au vipengele vingine ndani ya mstari. Mara nyingi hutumika kurekebisha nafasi ya wima ya vipengele kama picha, alama, au maandishi ndani ya mstari.

Kisha kamilisha kisanduku chako cha kuchagua kwa sifa ya `margin` yenye thamani ya `15px`.

# --hints--

Unapaswa kuongeza sifa ya `vertical-align` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign);
```

Unapaswa kuwa na sifa ya `vertical-align` yenye thamani ya `middle`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.verticalAlign, "middle");
```

Unapaswa kuongeza sifa ya `margin-right` kwa kichaguzi chako cha `input[type="checkbox"]`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight);
```

Unapaswa kuwa na sifa ya `margin-right` yenye thamani ya `15px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('input[type="checkbox"]')?.marginRight, "15px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

.settings-card {
  max-width: 250px;
  padding: 20px;
  border-radius: 10px;
  box-shadow: 0 2px 6px rgba(0, 0, 0, 0.2);
  margin: auto;
  text-align: left;
}

h1 {
  text-align: center;
}

label {
  display: block;
  cursor: pointer;
  margin: 8px auto;
}

input[type='checkbox'] {
  width: 20px;
  height: 20px;
  cursor: pointer;
  appearance: none;
  border: 2px solid #f1be32;
  border-radius: 4px;
  background-color: white;
  transition: all 0.3s;
--fcc-editable-region--
  
--fcc-editable-region--
}
```
