---
id: 68f29dbf9125476b4808d6fb
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Sasa ni wakati wa kupamba kontena la `settings-card`. Unda kichaguzi cha darasa la vitu kwa ajili ya `settings-card`.

Hapa ndipo utapanga muundo wote wa kontena lako. Weka `max-width` kuwa `250px` ili kufafanua ukubwa wa jumla wa kontena lako.

Kisha, weka `padding` kuwa `20px` ili maudhui yako yawe na nafasi kati yake na mpaka wa kontena.

Baada ya hapo, tengeneza kingo za mviringo kwa kuweka `border-radius` kuwa `10px`.

Kisha weka `box-shadow` yenye thamani ya `0 2px 6px rgba(0,0,0,0.2)`. Hii itaunda muonekano wa "kuinuliwa" kidogo ambao utaongeza kina kwa kontena.

# --hints--

Unapaswa kuwa na kichaguzi cha darasa la vitu katika css yako kwa `settings-card`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle(".settings-card"));
```

Unapaswa kuongeza sifa ya `max-width` kwa kichaguzi cha darasa la vitu cha `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth);
```

Unapaswa kuwa na sifa ya `max-width` yenye thamani ya `250px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.maxWidth, "250px");
```

Unapaswa kuongeza sifa ya `padding` kwa kichaguzi cha darasa la vitu cha `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding);
```

Unapaswa kuwa na sifa ya `padding` yenye thamani ya `20px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.padding, "20px");
```

Unapaswa kuongeza sifa ya `border-radius` kwa kichaguzi cha darasa la vitu chako cha `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius);
```

Unapaswa kuwa na sifa ya `border-radius` yenye thamani ya `10px`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.borderRadius, "10px");
```

Unapaswa kuongeza sifa ya `box-shadow` kwa kichaguzi cha darasa la vitu cha `.settings-card`.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow);
```

Unapaswa kuwa na `box-shadow` yenye thamani ya `0 2px 6px rgba(0,0,0,0.2)`.

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle(".settings-card")?.boxShadow, "rgba(0, 0, 0, 0.2) 0px 2px 6px");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Game Settings Panel</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <div class="settings-card">
      <h1>Game Settings</h1>
      <label> <input type="checkbox" />Sound Effects</label>
      <label> <input type="checkbox" />Background Music</label>
      <label> <input type="checkbox" />Hard Mode</label>
      <label> <input type="checkbox" />Haptic feedback</label>
    </div>
  </body>
</html>
```

```css
body {
  height: 100vh;
  background-color: #f0f0f0;
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--
```
