---
id: 6619240f46cec8e04d77e03a
title: 基本函数 练习A 
challengeType: 1
dashedName: top-basic-functions-exercise-a
---

# --description--

创建一个接收整数参数的函数。 当函数接收的整数小于 `10` 时，函数应该返回 `integer + 7` 。 如果整数大于或等于 `10`，函数应该返回给定的 `integer - 3`。

函数名应该是 `addOrSubtract` 。

# --hints--

你应该有一个名为 `addOrSubtract` 的函数。

```js
assert.isFunction(addOrSubtract);
```

你的函数应该接受一个整数作为参数。

```js
assert.match(addOrSubtract.toString(), /\s*addOrSubtract\(\s*\w+\s*\)/);
```

当函数接收的整数小于 10 时，你的函数应该返回 integer + 7 。

```js
assert.strictEqual(addOrSubtract(5), 12);
```

当函数接收的整数大于或等于 10 时，你的函数应该返回 integer - 3 。

```js
assert.strictEqual(addOrSubtract(10), 7);
```




# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function addOrSubtract(num) {
  if (num < 10) {
    return num + 7;
  } else {
    return num - 3;
  }
}
```
