---
id: 661e17c6068359c3ccf2f4d8
title: 基本函数 练习D
challengeType: 1
dashedName: top-basic-functions-exercise-d
---

# --description--

编写一个名为 `lastLetter` 的函数，该函数接收一个字符串作为参数，并返回字符串参数的最后一个字母。

# --hints--

你应该有一个名为 `lastLetter` 的函数。

```js
assert.isFunction(lastLetter);
```

你的函数应该接收一个字符串作为参数。 

```js
assert.match(lastLetter.toString(), /\s*lastLetter\(\s*\w+\s*\)/);
```

你应该返回字符串的最后一个字母。

```js
assert.strictEqual(lastLetter('Sem'), 'm');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function lastLetter(str) {
  return str[str.length - 1];
}
```
