---
id: 661e151f068359c3ccf2f4d7
title: 基本函数 练习C
challengeType: 1
dashedName: top-basic-functions-exercise-c
---

# --description--

编写一个名为 `capitalize` 的函数，接收一个字符串作为参数，并返回将接收字符串的首字母大写的新字符串。

# --hints--

你应该有一个名为 `capitalize` 的函数。

```js
assert.isFunction(capitalize);
```

你的函数应该接收一个字符串作为参数。 

```js
assert.match(capitalize.toString(), /\s*capitalize\(\s*\w+\s*\)/);
```

你的函数应该返回一个将接收字符串首字母大写的新字符串。

```js
assert.strictEqual(capitalize('sem'), 'Sem');
```


# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function capitalize(str) {
  return str[0].toUpperCase() + str.slice(1);
}
```
