---
id: 587d7db5367417b2b2512b96
title: Співпадіння букв алфавіту
challengeType: 1
forumTopicId: 301354
dashedName: match-letters-of-the-alphabet
---

# --description--

Ви бачили, як можна використовувати <dfn>множини символів</dfn>, щоб вказати групу символів для співпадіння, але це багато друкувати, коли потрібно співпасти з великим діапазоном символів (наприклад, з усіма буквами алфавіту). На щастя, є вбудована функція, яка робить це коротко і просто.

Всередині множини символів ви можете визначити діапазон символів для співпадіння, використовуючи дефіс: `-`.

Наприклад, щоб співпасти з літерами нижнього регістру `a` до `e`, ви б використали `[a-e]`.

```js
let catStr = "cat";
let batStr = "bat";
let matStr = "mat";
let bgRegex = /[a-e]at/;
catStr.match(bgRegex);
batStr.match(bgRegex);
matStr.match(bgRegex);
```

Послідовно три виклики `match` повернули б значення `["cat"]`, `["bat"]` і `null`.

# --instructions--

Співпасте з усіма буквами в рядку `quoteSample`.

**Примітка**: Обов’язково співпасте як з літерами верхнього, так і нижнього регістру.

# --hints--

Ваш регулярний вираз `alphabetRegex` має співпасти з 35 елементами.

```js
assert(result.length == 35);
```

Ваш регулярний вираз `alphabetRegex` має використовувати глобальний прапорець.

```js
assert(alphabetRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `alphabetRegex` має використовувати прапорець нечутливості до регістру.

```js
assert(alphabetRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /change/; // Change this line
let result = alphabetRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "The quick brown fox jumps over the lazy dog.";
let alphabetRegex = /[a-z]/gi; // Change this line
let result = quoteSample.match(alphabetRegex); // Change this line
```
