---
id: 587d7db4367417b2b2512b93
title: Знайти більше ніж перше співпадіння
challengeType: 1
forumTopicId: 301342
dashedName: find-more-than-the-first-match
---

# --description--

До цього моменту ви могли витягувати або шукати шаблон лише один раз.

```js
let testStr = "Repeat, Repeat, Repeat";
let ourRegex = /Repeat/;
testStr.match(ourRegex);
```

Тут `match` повернув би `["Repeat"]`.

Щоб шукати або витягувати шаблон більше одного разу, можна використовувати глобальний прапорець пошуку: `g`.

```js
let repeatRegex = /Repeat/g;
testStr.match(repeatRegex);
```

І тут `match` повертає значення `["Repeat", "Repeat", "Repeat"]`

# --instructions--

Використовуючи регулярний вираз `starRegex`, знайдіть і витягніть обидва слова `Twinkle` з рядка `twinkleStar`.

**Примітка**  
Ви можете використовувати кілька прапорців у своєму регулярному виразі, наприклад `/search/gi`

# --hints--

Ваш регулярний вираз `starRegex` має використовувати глобальний прапорець `g`

```js
assert(starRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `starRegex` має використовувати прапорець нечутливості до регістру `i`

```js
assert(starRegex.flags.match(/i/).length == 1);
```

Ваш збіг має відповідати обом входженням слова `Twinkle`

```js
assert(
  result.sort().join() ==
    twinkleStar
      .match(/twinkle/gi)
      .sort()
      .join()
);
```

Ваш збіг `result` має містити два елементи.

```js
assert(result.length == 2);
```

# --seed--

## --seed-contents--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /change/; // Change this line
let result = twinkleStar; // Change this line
```

# --solutions--

```js
let twinkleStar = "Twinkle, twinkle, little star";
let starRegex = /twinkle/gi;
let result = twinkleStar.match(starRegex);
```
