---
id: 587d7db5367417b2b2512b95
title: Зіставлення одного символу з кількома можливостями
challengeType: 1
forumTopicId: 301357
dashedName: match-single-character-with-multiple-possibilities
---

# --description--

Ви дізналися, як зіставляти буквальні шаблони (`/literal/`) і символи підстановки (`/./`). Це два крайні випадки регулярних виразів: один шукає точні співпадіння, а інший — все підряд. Існують варіанти, які є балансом між цими двома крайнощами.

Ви можете шукати буквальний шаблон з деякою гнучкістю за допомогою <dfn>класів символів</dfn>. Класи символів дозволяють визначити групу символів, які ви хочете зіставити, помістивши їх у квадратні дужки (`[` і `]`).

Наприклад, ви хочете зіставити `bag`, `big` і `bug`, але не `bog`. Ви можете створити регулярний вираз `/b[aiu]g/`, щоб це зробити. `[aiu]` — це клас символів, який буде зіставляти лише символи `a`, `i` або `u`.

```js
let bigStr = "big";
let bagStr = "bag";
let bugStr = "bug";
let bogStr = "bog";
let bgRegex = /b[aiu]g/;
bigStr.match(bgRegex);
bagStr.match(bgRegex);
bugStr.match(bgRegex);
bogStr.match(bgRegex);
```

Послідовно, чотири виклики `match` повернуть значення `["big"]`, `["bag"]`, `["bug"]` і `null`.

# --instructions--

Використайте клас символів із голосними (`a`, `e`, `i`, `o`, `u`) у вашому регулярному виразі `vowelRegex`, щоб знайти всі голосні у рядку `quoteSample`.

**Примітка:** Переконайтеся, що ви зіставляєте як великі, так і малі голосні.

# --hints--

Ви маєте знайти всі 25 голосних.

```js
assert(result.length == 25);
```

Ваш регулярний вираз `vowelRegex` має використовувати клас символів.

```js
assert(/\[.*\]/.test(vowelRegex.source));
```

Ваш регулярний вираз `vowelRegex` має використовувати глобальний прапорець.

```js
assert(vowelRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `vowelRegex` має використовувати прапорець нечутливості до регістру.

```js
assert(vowelRegex.flags.match(/i/).length == 1);
```

Ваш регулярний вираз не повинен зіставляти приголосні.

```js
assert(!/[b-df-hj-np-tv-z]/gi.test(result.join()));
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /change/; // Change this line
let result = vowelRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Beware of bugs in the above code; I have only proved it correct, not tried it.";
let vowelRegex = /[aeiou]/gi; // Change this line
let result = quoteSample.match(vowelRegex); // Change this line
```
