---
id: 587d7dba367417b2b2512ba8
title: Перевірка на все або нічого
challengeType: 1
forumTopicId: 301338
dashedName: check-for-all-or-none
---

# --description--

Іноді шаблони, які ви хочете знайти, можуть містити частини, які можуть існувати, а можуть і ні. Проте іноді важливо все одно перевірити їх наявність.

Ви можете вказати можливу наявність елемента за допомогою знаку питання, `?`. Це перевіряє нуль або один попередній елемент. Можна уявити цей символ як те, що попередній елемент є необов’язковим.

Наприклад, є невеликі відмінності між американською та британською англійською, і ви можете використати знак питання, щоб співпадали обидва варіанти написання.

```js
let american = "color";
let british = "colour";
let rainbowRegex= /colou?r/;
rainbowRegex.test(american);
rainbowRegex.test(british);
```

Обидва використання методу `test` повернули б `true`.

# --instructions--

Змініть регулярний вираз `favRegex`, щоб він співпадав як з американською англійською (`favorite`), так і з британською англійською (`favourite`) версією слова.

# --hints--

Ваш регулярний вираз має використовувати необов’язковий символ, `?`.

```js
favRegex.lastIndex = 0;
assert(favRegex.source.match(/\?/).length > 0);
```

Ваш регулярний вираз має співпадати з рядком `favorite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favorite'));
```

Ваш регулярний вираз має співпадати з рядком `favourite`

```js
favRegex.lastIndex = 0;
assert(favRegex.test('favourite'));
```

Ваш регулярний вираз не має співпадати з рядком `fav`

```js
favRegex.lastIndex = 0;
assert(!favRegex.test('fav'));
```

# --seed--

## --seed-contents--

```js
let favWord = "favorite";
let favRegex = /change/; // Change this line
let result = favRegex.test(favWord);
```

# --solutions--

```js
let favWord = "favorite";
let favRegex = /favou?r/;
let result = favRegex.test(favWord);
```
