---
id: 587d7db4367417b2b2512b90
title: Зіставлення буквального рядка з різними варіантами
challengeType: 1
forumTopicId: 301345
dashedName: match-a-literal-string-with-different-possibilities
---

# --description--

Використовуючи регулярні вирази, як `/coding/`, ви можете шукати шаблон `coding` в іншому рядку.

Це потужно для пошуку окремих рядків, але обмежено лише одним шаблоном. Ви можете шукати кілька шаблонів, використовуючи оператор `alternation` або `OR`: `|`.

Цей оператор зіставляє шаблони або перед ним, або після нього. Наприклад, якщо ви хочете знайти рядки `yes` або `no`, потрібний вам регулярний вираз — `/yes|no/`.

Ви також можете шукати більше ніж два шаблони. Це можна зробити, додавши більше шаблонів, розділених операторами `OR`, як у `/yes|no|maybe/`.

# --instructions--

Завершіть регулярний вираз `petRegex`, щоб знайти домашніх тварин `dog`, `cat`, `bird` або `fish`.

# --hints--

Ваш регулярний вираз `petRegex` має повертати `true` для рядка `John has a pet dog.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('John has a pet dog.'));
```

Ваш регулярний вираз `petRegex` має повертати `false` для рядка `Emma has a pet rock.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Emma has a pet rock.'));
```

Ваш регулярний вираз `petRegex` має повертати `true` для рядка `Emma has a pet bird.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Emma has a pet bird.'));
```

Ваш регулярний вираз `petRegex` має повертати `true` для рядка `Liz has a pet cat.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Liz has a pet cat.'));
```

Ваш регулярний вираз `petRegex` має повертати `false` для рядка `Kara has a pet dolphin.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Kara has a pet dolphin.'));
```

Ваш регулярний вираз `petRegex` має повертати `true` для рядка `Alice has a pet fish.`

```js
petRegex.lastIndex = 0;
assert(petRegex.test('Alice has a pet fish.'));
```

Ваш регулярний вираз `petRegex` має повертати `false` для рядка `Jimmy has a pet computer.`

```js
petRegex.lastIndex = 0;
assert(!petRegex.test('Jimmy has a pet computer.'));
```

# --seed--

## --seed-contents--

```js
let petString = "James has a pet cat.";
let petRegex = /change/; // Change this line
let result = petRegex.test(petString);
```

# --solutions--

```js
let petString = "James has a pet cat.";
let petRegex = /dog|cat|bird|fish/; // Change this line
let result = petRegex.test(petString);
```
