---
id: 587d7db9367417b2b2512ba4
title: "Знайти символи, що не є пробілами"
challengeType: 1
forumTopicId: 18210
dashedName: match-non-whitespace-characters
---

# --description--

Ви вивчили пошук пробілів за допомогою `\s` з нижнім регістром `s`. Ви також можете шукати все, крім пробілів.

Знайдіть символи, що не є пробілами, використовуючи `\S`, який є верхнім регістром `s`. Цей шаблон не співпадатиме з пробілами, поверненням каретки, табуляцією, формою подачі та символами нового рядка. Можна уявити, що це схоже на клас символів `[^ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let nonSpaceRegex = /\S/g;
whiteSpace.match(nonSpaceRegex).length;
```

Значення, яке поверне метод `.length`, буде `32`.

# --instructions--

Змініть регулярний вираз `countNonWhiteSpace`, щоб шукати кілька символів, що не є пробілами, у рядку.

# --hints--

Ваш регулярний вираз має використовувати глобальний прапорець.

```js
assert(countNonWhiteSpace.global);
```

Ваш регулярний вираз має використовувати скорочений клас символів `\S` для пошуку всіх символів, що не є пробілами.

```js
assert(/\\S/.test(countNonWhiteSpace.source));
```

Ваш регулярний вираз має знайти 35 символів, що не є пробілами, у рядку `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countNonWhiteSpace)
    .length == 35
);
```

Ваш регулярний вираз має знайти 23 символи, що не є пробілами, у рядку `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countNonWhiteSpace).length == 23);
```

Ваш регулярний вираз має знайти 21 символ, що не є пробілами, у рядку `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countNonWhiteSpace).length == 21);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /change/; // Change this line
let result = sample.match(countNonWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countNonWhiteSpace = /\S/g; // Change this line
let result = sample.match(countNonWhiteSpace);
```
