---
id: 587d7db7367417b2b2512b9f
title: Знайти всі літери та числа
challengeType: 1
forumTopicId: 301346
dashedName: match-all-letters-and-numbers
---

# --description--

Використовуючи класи символів, ви могли шукати всі літери алфавіту за допомогою `[a-z]`. Цей тип класу символів настільки поширений, що для нього існує скорочення, хоча воно також включає кілька додаткових символів.

Найближчим класом символів у JavaScript для пошуку алфавіту є `\w`. Це скорочення дорівнює `[A-Za-z0-9_]`. Цей клас символів відповідає великим і малим літерам, а також числам. Зверніть увагу, що цей клас символів також включає символ підкреслення (`_`).

```js
let longHand = /[A-Za-z0-9_]+/;
let shortHand = /\w+/;
let numbers = "42";
let varNames = "important_var";
longHand.test(numbers);
shortHand.test(numbers);
longHand.test(varNames);
shortHand.test(varNames);
```

Всі чотири ці виклики `test` повернули б `true`.

Ці скорочені класи символів також відомі як <dfn>скорочені класи символів</dfn>.

# --instructions--

Використайте скорочений клас символів `\w`, щоб порахувати кількість алфавітно-цифрових символів у різних цитатах і рядках.

# --hints--

Ваш регулярний вираз має використовувати глобальний прапорець.

```js
assert(alphabetRegexV2.global);
```

Ваш регулярний вираз має використовувати скорочений клас символів `\w`, щоб знайти всі алфавітно-цифрові символи.

```js
assert(/\\w/.test(alphabetRegexV2.source));
```

Ваш регулярний вираз має знайти 31 алфавітно-цифровий символ у рядку `The five boxing wizards jump quickly.`

```js
assert(
  'The five boxing wizards jump quickly.'.match(alphabetRegexV2).length === 31
);
```

Ваш регулярний вираз має знайти 32 алфавітно-цифрові символи у рядку `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(alphabetRegexV2).length ===
    32
);
```

Ваш регулярний вираз має знайти 30 алфавітно-цифрових символів у рядку `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(alphabetRegexV2).length === 30
);
```

Ваш регулярний вираз має знайти 36 алфавітно-цифрових символів у рядку `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(alphabetRegexV2)
    .length === 36
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /change/; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let alphabetRegexV2 = /\w/g; // Change this line
let result = quoteSample.match(alphabetRegexV2).length;
```
