---
id: 587d7db6367417b2b2512b99
title: "Збіг символів, які зустрічаються один або більше разів"
challengeType: 1
forumTopicId: 301350
dashedName: match-characters-that-occur-one-or-more-times
---

# --description--

Іноді потрібно знайти збіг символу (або групи символів), який з’являється один або більше разів підряд. Це означає, що він зустрічається принаймні один раз і може повторюватися.

Ви можете використати символ `+`, щоб перевірити, чи так є. Пам’ятайте, що символ або шаблон мають бути присутні послідовно. Тобто символ має повторюватися один за одним.

Наприклад, `/a+/g` знайде один збіг у `abc` і поверне `["a"]`. Через `+` він також знайде один збіг у `aabc` і поверне `["aa"]`.

Якщо ж перевіряти рядок `abab`, він знайде два збіги і поверне `["a", "a"]`, бо символи `a` не йдуть підряд — між ними є `b`. Нарешті, оскільки в рядку `a` немає `bcd`, збіг не буде знайдено.

# --instructions--

Ви хочете знайти збіги, коли літера `s` зустрічається один або більше разів у `Mississippi`. Напишіть регулярний вираз, який використовує знак `+`.

# --hints--

Ваш регулярний вираз `myRegex` має використовувати знак `+`, щоб знайти один або більше символів `s`.

```js
assert(/\+/.test(myRegex.source));
```

Ваш регулярний вираз `myRegex` має знайти 2 елементи.

```js
assert(result.length == 2);
```

Змінна `result` має бути масивом із двома збігами `ss`.

```js
assert(result[0] == 'ss' && result[1] == 'ss');
```

# --seed--

## --seed-contents--

```js
let difficultSpelling = "Mississippi";
let myRegex = /change/; // Change this line
let result = difficultSpelling.match(myRegex);
```

# --solutions--

```js
let difficultSpelling = "Mississippi";
let myRegex = /s+/g; // Change this line
let result = difficultSpelling.match(myRegex);
```
