---
id: 587d7db7367417b2b2512b9d
title: Збіг початкових шаблонів рядка
challengeType: 1
forumTopicId: 301349
dashedName: match-beginning-string-patterns
---

# --description--

Попередні завдання показали, що регулярні вирази можна використовувати для пошуку кількох збігів. Їх також застосовують для пошуку шаблонів у певних позиціях рядків.

У одному з попередніх завдань ви використовували символ карета (`^`) всередині множини символів, щоб створити заперечену множину у формі `[^thingsThatWillNotBeMatched]`. Поза межами множини символів карет використовується для пошуку шаблонів на початку рядків.

```js
let firstString = "Ricky is first and can be found.";
let firstRegex = /^Ricky/;
firstRegex.test(firstString);
let notFirst = "You can't find Ricky now.";
firstRegex.test(notFirst);
```

Перший виклик `test` поверне `true`, а другий — `false`.

# --instructions--

Використайте символ карета у регулярному виразі, щоб знайти `Cal` лише на початку рядка `rickyAndCal`.

# --hints--

Ваш регулярний вираз має шукати рядок `Cal` з великої літери.

```js
assert(calRegex.source == '^Cal');
```

Ваш регулярний вираз не повинен використовувати жодних прапорців.

```js
assert(calRegex.flags == '');
```

Ваш регулярний вираз має співпадати з рядком `Cal` на початку рядка.

```js
calRegex.lastIndex = 0;
assert(calRegex.test('Cal and Ricky both like racing.'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Cal` посередині рядка.

```js
calRegex.lastIndex = 0;
assert(!calRegex.test('Ricky and Cal both like racing.'));
```

# --seed--

## --seed-contents--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /change/; // Change this line
let result = calRegex.test(rickyAndCal);
```

# --solutions--

```js
let rickyAndCal = "Cal and Ricky both like racing.";
let calRegex = /^Cal/; // Change this line
let result = calRegex.test(rickyAndCal);
```
