---
id: 587d7db8367417b2b2512ba3
title: Знайти пробіли
challengeType: 1
forumTopicId: 301359
dashedName: match-whitespace
---

# --description--

Дотепер завдання охоплювали пошук літер алфавіту та чисел. Ви також можете знаходити пробіли або пропуски між літерами.

Ви можете шукати пробіли за допомогою `\s`, що є нижнім регістром `s`. Цей шаблон знаходить не лише пробіли, а й символи повернення каретки, табуляції, переводу форми та нового рядка. Можна уявити його як схожий на клас символів `[ \r\t\f\n\v]`.

```js
let whiteSpace = "Whitespace. Whitespace everywhere!"
let spaceRegex = /\s/g;
whiteSpace.match(spaceRegex);
```

Цей виклик `match` поверне `[" ", " "]`.
# --instructions--

Змініть регулярний вираз `countWhiteSpace`, щоб він шукав кілька пробілів у рядку.

# --hints--

Ваш регулярний вираз має використовувати глобальний прапорець.

```js
assert(countWhiteSpace.global);
```

Ваш регулярний вираз має використовувати скорочений символ `\s` для пошуку всіх пробілів.

```js
assert(/\\s/.test(countWhiteSpace.source));
```

Ваш регулярний вираз має знайти вісім пробілів у рядку `Men are from Mars and women are from Venus.`

```js
assert(
  'Men are from Mars and women are from Venus.'.match(countWhiteSpace).length ==
    8
);
```

Ваш регулярний вираз має знайти три пробіли у рядку `Space: the final frontier.`

```js
assert('Space: the final frontier.'.match(countWhiteSpace).length == 3);
```

Ваш регулярний вираз не має знаходити пробілів у рядку `MindYourPersonalSpace`

```js
assert('MindYourPersonalSpace'.match(countWhiteSpace) == null);
```

# --seed--

## --seed-contents--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /change/; // Change this line
let result = sample.match(countWhiteSpace);
```

# --solutions--

```js
let sample = "Whitespace is important in separating words";
let countWhiteSpace = /\s/g;
let result = sample.match(countWhiteSpace);
```
