---
id: 587d7db3367417b2b2512b8e
title: Використання методу test
challengeType: 1
forumTopicId: 301369
dashedName: using-the-test-method
---

# --description--

Регулярні вирази використовують у мовах програмування, щоб знаходити частини рядків. Ви створюєте шаблони, які допомагають вам робити це знаходження.

Якщо ви хочете знайти слово `the` у рядку `The dog chased the cat`, ви можете використати такий регулярний вираз: `/the/`. Зверніть увагу, що лапки не потрібні всередині регулярного виразу.

У JavaScript є кілька способів використання регулярних виразів. Один зі способів перевірити регулярний вираз — це використати метод `.test()`. Метод `.test()` приймає регулярний вираз, застосовує його до рядка (який розміщується всередині круглих дужок) і повертає `true` або `false`, залежно від того, чи знаходить ваш шаблон щось.

```js
let testStr = "freeCodeCamp";
let testRegex = /Code/;
testRegex.test(testStr);
```

Метод `test` тут повертає `true`.

# --instructions--

Застосуйте регулярний вираз `myRegex` до рядка `myString` за допомогою методу `.test()`.

# --hints--

Ви повинні використати `.test()` для перевірки регулярного виразу.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Ваш результат має повернути `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex; // Change this line
```

# --solutions--

```js
let myString = "Hello, World!";
let myRegex = /Hello/;
let result = myRegex.test(myString); // Change this line
```
