---
id: 587d7db3367417b2b2512b8f
title: Збіг буквальних рядків
challengeType: 1
forumTopicId: 301355
dashedName: match-literal-strings
---

# --description--

У попередньому завданні ви шукали слово `Hello` за допомогою регулярного виразу `/Hello/`. Цей регулярний вираз шукав буквальний збіг рядка `Hello`. Ось ще один приклад пошуку буквального збігу рядка `Kevin`:

```js
let testStr = "Hello, my name is Kevin.";
let testRegex = /Kevin/;
testRegex.test(testStr);
```

Цей виклик `test` поверне `true`.

Жодні інші форми `Kevin` не співпадуть. Наприклад, регулярний вираз `/Kevin/` не співпаде з `kevin` або `KEVIN`.

```js
let wrongRegex = /kevin/;
wrongRegex.test(testStr);
```

Цей виклик `test` поверне `false`.

У майбутньому завданні буде показано, як також співпадати з іншими формами.

# --instructions--

Завершіть регулярний вираз `waldoRegex`, щоб знайти `"Waldo"` у рядку `waldoIsHiding` з буквальним збігом.

# --hints--

Ваш регулярний вираз `waldoRegex` має знайти рядок `Waldo`

```js
waldoRegex.lastIndex = 0;
assert(waldoRegex.test(waldoIsHiding));
```

Ваш регулярний вираз `waldoRegex` не повинен шукати нічого іншого.

```js
waldoRegex.lastIndex = 0;
assert(!waldoRegex.test('Somewhere is hiding in this text.'));
```

Ви маєте виконати буквальний збіг рядка за допомогою регулярного виразу.

```js
assert(!/\/.*\/i/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /search/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```

# --solutions--

```js
let waldoIsHiding = "Somewhere Waldo is hiding in this text.";
let waldoRegex = /Waldo/; // Change this line
let result = waldoRegex.test(waldoIsHiding);
```
