---
id: 587d7db4367417b2b2512b92
title: витягти співпадіння
challengeType: 1
forumTopicId: 301340
dashedName: extract-matches
---

# --description--

До цього моменту ви лише перевіряли, чи існує шаблон у рядку. Ви також можете витягти фактичні співпадіння, які знайшли, за допомогою методу `.match()`.

Щоб використати метод `.match()`, застосуйте цей метод до рядка і передайте регулярний вираз у круглих дужках.

Ось приклад:

```js
"Hello, World!".match(/Hello/);
let ourStr = "Regular expressions";
let ourRegex = /expressions/;
ourStr.match(ourRegex);
```

Тут перший `match` поверне `["Hello"]`, а другий поверне `["expressions"]`.

Зверніть увагу, що синтаксис `.match` є «протилежним» до методу `.test`, який ви використовували досі:

```js
'string'.match(/regex/);
/regex/.test('string');
```

# --instructions--

Застосуйте метод `.match()`, щоб витягти рядок `coding`.

# --hints--

`result` має містити рядок `coding`

```js
assert(result.join() === 'coding');
```

Ваш регулярний вираз `codingRegex` має шукати рядок `coding`

```js
assert(codingRegex.source === 'coding');
```

Ви маєте використати метод `.match()`.

```js
assert(__helpers.removeJSComments(code).match(/\.match\(.*\)/));
```

# --seed--

## --seed-contents--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /change/; // Change this line
let result = extractStr; // Change this line
```

# --solutions--

```js
let extractStr = "Extract the word 'coding' from this string.";
let codingRegex = /coding/; // Change this line
let result = extractStr.match(codingRegex); // Change this line
```
