---
id: 587d7db4367417b2b2512b91
title: Ігнорувати регістр під час співпадіння
challengeType: 1
forumTopicId: 301344
dashedName: ignore-case-while-matching
---

# --description--

До цього моменту ви розглядали регулярні вирази для точного співпадіння рядків. Але іноді ви можете захотіти також враховувати відмінності регістру.

Регістр (або іноді літературний регістр) — це різниця між великими та малими літерами. Прикладами великих літер є `A`, `B` і `C`. Прикладами малих літер є `a`, `b` і `c`.

Ви можете співпадати з обома регістрами, використовуючи так званий прапорець. Існують інші прапорці, але тут ви зосередитесь на прапорці, який ігнорує регістр — прапорці `i`. Ви можете використати його, додавши в кінці регулярного виразу. Приклад використання цього прапорця — `/ignorecase/i`. Цей регулярний вираз може співпадати з рядками `ignorecase`, `igNoreCase` і `IgnoreCase`.

# --instructions--

Напишіть регулярний вираз `fccRegex`, який співпадатиме з `freeCodeCamp`, незалежно від регістру. Ваш регулярний вираз не повинен співпадати з абревіатурами або варіаціями з пробілами.

# --hints--

Ваш регулярний вираз повинен співпадати з рядком `freeCodeCamp`

```js
fccRegex.lastIndex = 0; 
assert(fccRegex.test('freeCodeCamp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FreeCodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodeCamp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FreecodeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreecodeCamp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FreeCodecamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCodecamp'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Free Code Camp`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('Free Code Camp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FreeCOdeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FreeCOdeCamp'));
```

Ваш регулярний вираз не повинен співпадати з рядком `FCC`

```js
fccRegex.lastIndex = 0;
assert(!fccRegex.test('FCC'));
```

Ваш регулярний вираз повинен співпадати з рядком `FrEeCoDeCamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCoDeCamp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FrEeCodECamp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FrEeCodECamp'));
```

Ваш регулярний вираз повинен співпадати з рядком `FReeCodeCAmp`

```js
fccRegex.lastIndex = 0;
assert(fccRegex.test('FReeCodeCAmp'));
```

# --seed--

## --seed-contents--

```js
let myString = "freeCodeCamp";
let fccRegex = /change/; // Change this line
let result = fccRegex.test(myString);
```

# --solutions--

```js
let myString = "freeCodeCamp";
let fccRegex = /freecodecamp/i; // Change this line
let result = fccRegex.test(myString);
```
