---
id: 587d7db8367417b2b2512ba1
title: Знайти всі нечислові символи
challengeType: 1
forumTopicId: 301347
dashedName: match-all-non-numbers
---

# --description--

У попередньому завданні показано, як шукати цифри за допомогою скорочення `\d` з нижнім регістром `d`. Ви також можете шукати нецифри, використовуючи подібне скорочення, але з верхнім регістром `D`.

Скорочення для пошуку нечислових символів — `\D`. Воно дорівнює класу символів `[^0-9]`, який шукає один символ, що не є числом від нуля до дев’яти.

# --instructions--

Використайте скорочений клас символів для нецифр `\D`, щоб порахувати, скільки нецифр у назвах фільмів.

# --hints--

Ваш регулярний вираз має використовувати скорочений символ для пошуку нечислових символів

```js
assert(/\\D/.test(noNumRegex.source));
```

Ваш регулярний вираз має використовувати глобальний прапорець.

```js
assert(noNumRegex.global);
```

Ваш регулярний вираз не має знаходити нецифр у рядку `9`.

```js
assert('9'.match(noNumRegex) == null);
```

Ваш регулярний вираз має знайти 6 нецифр у рядку `Catch 22`.

```js
assert('Catch 22'.match(noNumRegex).length == 6);
```

Ваш регулярний вираз має знайти 11 нецифр у рядку `101 Dalmatians`.

```js
assert('101 Dalmatians'.match(noNumRegex).length == 11);
```

Ваш регулярний вираз має знайти 15 нецифр у рядку `One, Two, Three`.

```js
assert('One, Two, Three'.match(noNumRegex).length == 15);
```

Ваш регулярний вираз має знайти 12 нецифр у рядку `21 Jump Street`.

```js
assert('21 Jump Street'.match(noNumRegex).length == 12);
```

Ваш регулярний вираз має знайти 17 нецифр у рядку `2001: A Space Odyssey`.

```js
assert('2001: A Space Odyssey'.match(noNumRegex).length == 17);
```

# --seed--

## --seed-contents--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /change/; // Change this line
let result = movieName.match(noNumRegex).length;
```

# --solutions--

```js
let movieName = "2001: A Space Odyssey";
let noNumRegex = /\D/g; // Change this line
let result = movieName.match(noNumRegex).length;
```
