---
id: 587d7db5367417b2b2512b97
title: Співпадіння чисел і літер алфавіту
challengeType: 1
forumTopicId: 301356
dashedName: match-numbers-and-letters-of-the-alphabet
---

# --description--

Використання дефіса (`-`) для співпадіння діапазону символів не обмежується лише літерами. Він також працює для співпадіння діапазону чисел.

Наприклад, `/[0-5]/` співпадає з будь-яким числом між `0` і `5`, включно з `0` і `5`.

Також можливо поєднати діапазон літер і чисел у одному наборі символів.

```js
let jennyStr = "Jenny8675309";
let myRegex = /[a-z0-9]/ig;
jennyStr.match(myRegex);
```

# --instructions--

Створіть один регулярний вираз, який співпадає з діапазоном літер між `h` і `s`, а також з діапазоном чисел між `2` і `6`. Не забудьте включити відповідні прапорці у регулярний вираз.

# --hints--

Ваш регулярний вираз `myRegex` має співпадати з 17 елементами.

```js
assert(result.length == 17);
```

Ваш регулярний вираз `myRegex` має використовувати глобальний прапорець.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `myRegex` має використовувати прапорець нечутливості до регістру.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "Blueberry 3.141592653s are delicious.";
let myRegex = /[h-s2-6]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
