---
id: 587d7db9367417b2b2512ba6
title: Вказати лише нижню межу кількості співпадінь
challengeType: 1
forumTopicId: 301366
dashedName: specify-only-the-lower-number-of-matches
---

# --description--

Ви можете вказати нижню та верхню межу кількості шаблонів за допомогою кількісних специфікаторів у фігурних дужках. Іноді потрібно вказати лише нижню межу кількості шаблонів без верхньої межі.

Щоб вказати лише нижню межу кількості шаблонів, залиште перше число, після якого йде кома.

Наприклад, щоб співпадало лише з рядком `hah`, у якому літера `a` з’являється щонайменше `3` разів, ваш регулярний вираз буде `/ha{3,}h/`.

```js
let A4 = "haaaah";
let A2 = "haah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleA = /ha{3,}h/;
multipleA.test(A4);
multipleA.test(A2);
multipleA.test(A100);
```

Послідовно три виклики `test` повернуть `true`, `false` і `true`.

# --instructions--

Змініть регулярний вираз `haRegex`, щоб він співпадав зі словом `Hazzah` лише тоді, коли в ньому є чотири або більше літер `z`.

# --hints--

Ваш регулярний вираз має використовувати фігурні дужки.

```js
assert(haRegex.source.match(/{.*?}/).length > 0);
```

Ваш регулярний вираз не повинен співпадати з рядком `Hazzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzah'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Hazzzah`

```js
haRegex.lastIndex = 0;
assert(!haRegex.test('Hazzzah'));
```

Ваш регулярний вираз має співпадати з рядком `Hazzzzah`

```js
assert('Hazzzzah'.match(haRegex)[0].length === 8);
```

Ваш регулярний вираз має співпадати з рядком `Hazzzzzah`

```js
assert('Hazzzzzah'.match(haRegex)[0].length === 9);
```

Ваш регулярний вираз має співпадати з рядком `Hazzzzzzah`

```js
assert('Hazzzzzzah'.match(haRegex)[0].length === 10);
```

Ваш регулярний вираз має співпадати з рядком `Hazzah` з 30 `z` у ньому.

```js
assert('Hazzzzzzzzzzzzzzzzzzzzzzzzzzzzzzah'.match(haRegex)[0].length === 34);
```

# --seed--

## --seed-contents--

```js
let haStr = "Hazzzzah";
let haRegex = /change/; // Change this line
let result = haRegex.test(haStr);
```

# --solutions--

```js
let haStr = "Hazzzzah";
let haRegex = /Haz{4,}ah/; // Change this line
let result = haRegex.test(haStr);
```
