---
id: 587d7db6367417b2b2512b9a
title: "Зіставлення символів, які зустрічаються нуль або більше разів"
challengeType: 1
forumTopicId: 301351
dashedName: match-characters-that-occur-zero-or-more-times
---

# --description--

У попередньому завданні використовувався знак плюс `+` для пошуку символів, які зустрічаються один або більше разів. Існує також варіант, який відповідає символам, що зустрічаються нуль або більше разів.

Символом для цього є зірочка або астериск: `*`.

```js
let soccerWord = "gooooooooal!";
let gPhrase = "gut feeling";
let oPhrase = "over the moon";
let goRegex = /go*/;
soccerWord.match(goRegex);
gPhrase.match(goRegex);
oPhrase.match(goRegex);
```

Послідовно три виклики `match` повернули б значення `["goooooooo"]`, `["g"]` та `null`.

# --instructions--

Створіть регулярний вираз `chewieRegex`, який використовує символ `*` для зіставлення великої літери `A`, за якою йдуть нуль або більше літер у нижньому регістрі `a` у рядку `chewieQuote`. Ваш регулярний вираз не потребує прапорців або класів символів і не повинен відповідати жодним іншим цитатам.

# --hints--

Ваш регулярний вираз `chewieRegex` має використовувати символ `*` для зіставлення нуля або більше символів `a`.

```js
assert(/\*/.test(chewieRegex.source));
```

Ваш регулярний вираз має відповідати рядку `A` у `chewieQuote`.

```js
assert(result[0][0] === 'A');
```

Ваш регулярний вираз має відповідати рядку `Aaaaaaaaaaaaaaaa` у `chewieQuote`.

```js
assert(result[0] === 'Aaaaaaaaaaaaaaaa');
```

Ваш регулярний вираз `chewieRegex` має відповідати 16 символам у `chewieQuote`.

```js
assert(result[0].length === 16);
```

Ваш регулярний вираз не повинен відповідати жодним символам у рядку `He made a fair move. Screaming about it can't help you.`

```js
assert(
  !"He made a fair move. Screaming about it can't help you.".match(chewieRegex)
);
```

Ваш регулярний вираз не повинен відповідати жодним символам у рядку `Let him have it. It's not wise to upset a Wookiee.`

```js
assert(
  !"Let him have it. It's not wise to upset a Wookiee.".match(chewieRegex)
);
```

# --seed--

## --seed-contents--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
// Only change code below this line
let chewieRegex = /change/; // Change this line
// Only change code above this line

let result = chewieQuote.match(chewieRegex);
```

# --solutions--

```js
const chewieQuote = "Aaaaaaaaaaaaaaaarrrgh!";
let chewieRegex = /Aa*/;
let result = chewieQuote.match(chewieRegex);
```
