---
id: 587d7db7367417b2b2512b9c
title: Знайти одного або більше злочинців під час полювання
challengeType: 1
forumTopicId: 301343
dashedName: find-one-or-more-criminals-in-a-hunt
---

# --description--

Час зробити паузу і перевірити свої нові навички написання регулярних виразів. Група злочинців втекла з в’язниці і втекла, але ви не знаєте, скільки їх. Однак ви знаєте, що вони тримаються разом, коли знаходяться серед інших людей. Ви відповідаєте за те, щоб знайти всіх злочинців одразу.

Ось приклад, щоб повторити, як це зробити:

Регулярний вираз `/z+/` співпадає з літерою `z`, коли вона з’являється один або більше разів підряд. Він знайде співпадіння у всіх наступних рядках:

```js
"z"
"zzzzzz"
"ABCzzzz"
"zzzzABC"
"abczzzzzzzzzzzzzzzzzzzzzabc"
```

Але він не знайде співпадінь у наступних рядках, оскільки там немає літер `z`:

```js
""
"ABC"
"abcabc"
```

# --instructions--

Напишіть жадібний регулярний вираз, який знаходить одного або більше злочинців у групі інших людей. Злочинець представлений великою літерою `C`.

# --hints--

Ваш регулярний вираз має співпадати з одним злочинцем (`C`) у рядку `C`

```js
assert('C'.match(reCriminals) && 'C'.match(reCriminals)[0] == 'C');
```

Ваш регулярний вираз має співпадати з двома злочинцями (`CC`) у рядку `CC`

```js
assert('CC'.match(reCriminals) && 'CC'.match(reCriminals)[0] == 'CC');
```

Ваш регулярний вираз має співпадати з трьома злочинцями (`CCC`) у рядку `P1P5P4CCCcP2P6P3`.

```js
assert(
  'P1P5P4CCCcP2P6P3'.match(reCriminals) &&
    'P1P5P4CCCcP2P6P3'.match(reCriminals)[0] == 'CCC'
);
```

Ваш регулярний вираз має співпадати з п’ятьма злочинцями (`CCCCC`) у рядку `P6P2P7P4P5CCCCCP3P1`

```js
assert(
  'P6P2P7P4P5CCCCCP3P1'.match(reCriminals) &&
    'P6P2P7P4P5CCCCCP3P1'.match(reCriminals)[0] == 'CCCCC'
);
```

Ваш регулярний вираз не має співпадати з жодним злочинцем у порожньому рядку `""`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test(''));
```

Ваш регулярний вираз не має співпадати з жодним злочинцем у рядку `P1P2P3`

```js
reCriminals.lastIndex = 0;
assert(!reCriminals.test('P1P2P3'));
```

Ваш регулярний вираз має співпадати з п’ятдесятьма злочинцями (`CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC`) у рядку `P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3`.

```js
assert(
  'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
    reCriminals
  ) &&
    'P2P1P5P4CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCP3'.match(
      reCriminals
    )[0] == 'CCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCCC'
);
```

# --seed--

## --seed-contents--

```js
let reCriminals = /./; // Change this line
```

# --solutions--

```js
let reCriminals = /C+/; // Change this line
```
