---
id: 587d7db8367417b2b2512ba2
title: Обмеження можливих імен користувачів
challengeType: 1
forumTopicId: 301363
dashedName: restrict-possible-usernames
---

# --description--

Імена користувачів використовуються скрізь в інтернеті. Вони дають користувачам унікальну ідентичність на улюблених сайтах.

Вам потрібно перевірити всі імена користувачів у базі даних. Ось кілька простих правил, яких користувачі мають дотримуватися при створенні свого імені користувача.

1) Імена користувачів можуть містити лише алфавітно-цифрові символи.

2) Єдині числа в імені користувача мають бути в кінці. Там може бути нуль або більше чисел. Ім’я користувача не може починатися з числа.

3) Літери імені користувача можуть бути в нижньому або верхньому регістрі.

4) Імена користувачів мають бути не коротшими за два символи. Ім’я з двох символів може містити лише літери алфавіту.

# --instructions--

Змініть регулярний вираз `userCheck` так, щоб він відповідав наведеним вище обмеженням.

# --hints--

Ваш регулярний вираз має співпадати з рядком `JACK`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('JACK'));
```

Ваш регулярний вираз не має співпадати з рядком `J`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J'));
```

Ваш регулярний вираз має співпадати з рядком `Jo`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Jo'));
```

Ваш регулярний вираз має співпадати з рядком `Oceans11`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Oceans11'));
```

Ваш регулярний вираз має співпадати з рядком `RegexGuru`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('RegexGuru'));
```

Ваш регулярний вираз не має співпадати з рядком `007`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('007'));
```

Ваш регулярний вираз не має співпадати з рядком `9`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('9'));
```

Ваш регулярний вираз не має співпадати з рядком `A1`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('A1'));
```

Ваш регулярний вираз не має співпадати з рядком `BadUs3rnam3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('BadUs3rnam3'));
```

Ваш регулярний вираз має співпадати з рядком `Z97`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('Z97'));
```

Ваш регулярний вираз не має співпадати з рядком `c57bT3`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('c57bT3'));
```

Ваш регулярний вираз має співпадати з рядком `AB1`

```js
userCheck.lastIndex = 0;
assert(userCheck.test('AB1'));
```

Ваш регулярний вираз не має співпадати з рядком `J%4`

```js
userCheck.lastIndex = 0;
assert(!userCheck.test('J%4'))
```

# --seed--

## --seed-contents--

```js
let username = "JackOfAllTrades";
let userCheck = /change/; // Change this line
let result = userCheck.test(username);
```

# --solutions--

```js
let username = "JackOfAllTrades";
const userCheck = /^[a-z]([0-9]{2,}|[a-z]+\d*)$/i;
let result = userCheck.test(username);
```
