---
id: 587d7db7367417b2b2512b9e
title: Збіг шаблонів у кінці рядка
challengeType: 1
forumTopicId: 301352
dashedName: match-ending-string-patterns
---

# --description--

У попередньому завданні ви навчились використовувати символ карета для пошуку шаблонів на початку рядків. Існує також спосіб шукати шаблони в кінці рядків.

Ви можете шукати в кінці рядків, використовуючи символ долара `$` у кінці регулярного виразу.

```js
let theEnding = "This is a never ending story";
let storyRegex = /story$/;
storyRegex.test(theEnding);
let noEnding = "Sometimes a story will have to end";
storyRegex.test(noEnding);
```

Перший виклик `test` поверне `true`, а другий поверне `false`.

# --instructions--

Використайте символ-якір (`$`) для пошуку рядка `caboose` у кінці рядка `caboose`.

# --hints--

Ви повинні шукати `caboose` із символом долара `$` як якір у вашому регулярному виразі.

```js
assert(lastRegex.source == 'caboose$');
```

Ваш регулярний вираз не повинен містити прапорців.

```js
assert(lastRegex.flags == '');
```

Ви повинні співставити `caboose` у кінці рядка `The last car on a train is the caboose`

```js
lastRegex.lastIndex = 0;
assert(lastRegex.test('The last car on a train is the caboose'));
```

# --seed--

## --seed-contents--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /change/; // Change this line
let result = lastRegex.test(caboose);
```

# --solutions--

```js
let caboose = "The last car on a train is the caboose";
let lastRegex = /caboose$/; // Change this line
let result = lastRegex.test(caboose);
```
