---
id: 587d7dbb367417b2b2512bab
title: Використання груп захоплення для пошуку та заміни
challengeType: 1
forumTopicId: 301368
dashedName: use-capture-groups-to-search-and-replace
---

# --description--

Пошук дуже корисний. Однак ви можете зробити пошук ще потужнішим, якщо він також змінює (або замінює) текст, який ви знаходите.

Ви можете шукати та замінювати текст у рядку, використовуючи `.replace()` на рядку. Вхідними даними для `.replace()` є спочатку регулярний вираз, який ви хочете знайти. Другий параметр — рядок для заміни знайденого або функція для виконання дії.

```js
let wrongText = "The sky is silver.";
let silverRegex = /silver/;
wrongText.replace(silverRegex, "blue");
```

Виклик `replace` поверне рядок `The sky is blue.`.

Ви також можете отримати доступ до груп захоплення у рядку заміни за допомогою знаків долара (`$`).

```js
"Code Camp".replace(/(\w+)\s(\w+)/, '$2 $1');
```

Виклик `replace` поверне рядок `Camp Code`.

# --instructions--

Напишіть регулярний вираз `fixRegex` з трьома групами захоплення, який шукатиме кожне слово у рядку `one two three`. Потім оновіть змінну `replaceText`, щоб замінити `one two three` на рядок `three two one` і призначте результат змінній `result`. Переконайтеся, що ви використовуєте групи захоплення у рядку заміни за допомогою синтаксису зі знаком долара (`$`).

# --hints--

Ви повинні використовувати `.replace()` для пошуку та заміни.

```js
assert(__helpers.removeJSComments(code).match(/\.replace\(.*\)/));
```

Ваш регулярний вираз має змінити рядок `one two three` на рядок `three two one`

```js
assert(result === 'three two one');
```

Ви не повинні змінювати останній рядок.

```js
assert(__helpers.removeJSComments(code).match(/result\s*=\s*str\.replace\(.*?\)/));
```

`fixRegex` має містити щонайменше три групи захоплення.

```js
assert(new RegExp(fixRegex.source + '|').exec('').length - 1 >= 3);
```

`replaceText` має використовувати підрядки в дужках (тобто n-та підрядкова група в дужках, $n, відповідає n-й групі захоплення).

```js
{
  const re = /(\$\d{1,2})+(?:[\D]|\b)/g;
  assert(replaceText.match(re).length >= 3);
}
```

# --seed--

## --seed-contents--

```js
let str = "one two three";
let fixRegex = /change/; // Change this line
let replaceText = ""; // Change this line
let result = str.replace(fixRegex, replaceText);
```

# --solutions--

```js
let str = "one two three";
let fixRegex = /(\w+) (\w+) (\w+)/g; // Change this line
let replaceText = "$3 $2 $1"; // Change this line
let result = str.replace(fixRegex, replaceText);
```
