---
id: 587d7dbb367417b2b2512baa
title: Повторне використання шаблонів за допомогою груп захоплення
challengeType: 1
forumTopicId: 301364
dashedName: reuse-patterns-using-capture-groups
---

# --description--

Припустімо, ви хочете знайти слово, яке повторюється кілька разів, як показано нижче.

```js
let repeatStr = "row row row your boat";
```

Ви могли б використати `/row row row/`, але що, якщо ви не знаєте конкретне слово, яке повторюється? <dfn>Групи захоплення</dfn> можна використати, щоб знайти повторювані підрядки.

Групи захоплення створюються шляхом взяття регулярного виразу, який потрібно захопити, у круглі дужки. У цьому випадку мета — захопити слово, що складається з алфавітно-цифрових символів, тому група захоплення буде `\w+`, взята у круглі дужки: `/(\w+)/`.

Підрядок, який відповідає групі, зберігається у тимчасову «змінну», до якої можна звернутися в межах того самого регулярного виразу за допомогою зворотної скісної риски та номера групи захоплення (наприклад, `\1`). Групи захоплення автоматично нумеруються за позицією їх відкриваючих круглих дужок (зліва направо), починаючи з 1.

Приклад нижче знаходить слово, яке зустрічається тричі, розділене пробілами:

```js
let repeatRegex = /(\w+) \1 \1/;
repeatRegex.test(repeatStr); // Returns true
repeatStr.match(repeatRegex); // Returns ["row row row", "row"]
```

Використання методу `.match()` на рядку поверне масив із знайденим підрядком разом із його захопленими групами.


# --instructions--

Використайте групи захоплення у `reRegex`, щоб знайти рядок, який складається лише з одного й того самого числа, повтореного рівно тричі, розділеного одиночними пробілами.

# --hints--

Ваш регулярний вираз має використовувати скорочений клас символів для цифр.

```js
assert(reRegex.source.match(/\\d/));
```

Ваш регулярний вираз має повторно використовувати групу захоплення двічі.

```js
assert(reRegex.source.match(/\\1|\\2/g).length >= 2);
```

Ваш регулярний вираз має відповідати рядку `42 42 42`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('42 42 42'));
```

Ваш регулярний вираз має відповідати рядку `100 100 100`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('100 100 100'));
```

Ваш регулярний вираз не має відповідати рядку `42 42 42 42`.

```js
assert.equal('42 42 42 42'.match(reRegex.source), null);
```

Ваш регулярний вираз не має відповідати рядку `42 42`.

```js
assert.equal('42 42'.match(reRegex.source), null);
```

Ваш регулярний вираз не має відповідати рядку `101 102 103`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('101 102 103'));
```

Ваш регулярний вираз не має відповідати рядку `1 2 3`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('1 2 3'));
```

Ваш регулярний вираз має відповідати рядку `10 10 10`.

```js
reRegex.lastIndex = 0;
assert(reRegex.test('10 10 10'));
```

Ваш регулярний вираз не має відповідати рядку `42\t42\t42`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42\t42\t42'));
```

Ваш регулярний вираз не має відповідати рядку `42  42  42`.

```js
reRegex.lastIndex = 0;
assert(!reRegex.test('42  42  42'));
```

# --seed--

## --seed-contents--

```js
let repeatNum = "42 42 42";
let reRegex = /change/; // Change this line
let result = reRegex.test(repeatNum);
```

# --solutions--

```js
let repeatNum = "42 42 42";
let reRegex = /^(\d+) \1 \1$/;
let result = reRegex.test(repeatNum);
```
