---
id: 587d7dbb367417b2b2512bac
title: Видалити пробіли на початку та в кінці
challengeType: 1
forumTopicId: 301362
dashedName: remove-whitespace-from-start-and-end
---

# --description--

Іноді пробільні символи навколо рядків небажані, але вони присутні. Типова обробка рядків полягає у видаленні пробілів на початку та в кінці рядка.

# --instructions--

Напишіть регулярний вираз і використайте відповідні методи рядка, щоб видалити пробіли на початку та в кінці рядків.

**Примітка:** Метод `String.prototype.trim()` міг би тут спрацювати, але вам потрібно виконати це завдання, використовуючи регулярні вирази.

# --hints--

`result` має бути рівним рядку `Hello, World!`

```js
assert(result === 'Hello, World!');
```

Ваше рішення не повинно використовувати метод `String.prototype.trim()`.

```js
assert(!__helpers.removeJSComments(code).match(/\.?[\s\S]*?trim/));
```

Змінна `result` не повинна безпосередньо присвоюватися рядком

```js
assert(!__helpers.removeJSComments(code).match(/result\s*=\s*["'`].*?["'`]/));
```

Значення змінної `hello` не повинно змінюватися.

```js
assert(hello === '   Hello, World!  ');
```

# --seed--

## --seed-contents--

```js
let hello = "   Hello, World!  ";
let wsRegex = /change/; // Change this line
let result = hello; // Change this line
```

# --solutions--

```js
let hello = "   Hello, World!  ";
let wsRegex = /^(\s+)(.+[^\s])(\s+)$/;
let result = hello.replace(wsRegex, '$2');
```
