---
id: 587d7db5367417b2b2512b94
title: Збіг будь-чого за допомогою символу підстановки
challengeType: 1
forumTopicId: 301348
dashedName: match-anything-with-wildcard-period
---

# --description--

Іноді ви не будете (або не потрібно) знати точні символи у ваших шаблонах. Перебирати всі слова, які підходять, наприклад, для помилки в написанні, зайняло б багато часу. На щастя, ви можете заощадити час, використовуючи символ підстановки: `.`

Символ підстановки `.` відповідатиме будь-якому одному символу. Символ підстановки також називають `dot` та `period`. Ви можете використовувати символ підстановки так само, як і будь-який інший символ у регулярному виразі. Наприклад, якщо ви хочете знайти збіг для `hug`, `huh`, `hut` і `hum`, ви можете використати регулярний вираз `/hu./`, щоб знайти всі чотири слова.

```js
let humStr = "I'll hum a song";
let hugStr = "Bear hug";
let huRegex = /hu./;
huRegex.test(humStr);
huRegex.test(hugStr);
```

Обидва ці виклики `test` повернуть `true`.

# --instructions--

Завершіть регулярний вираз `unRegex`, щоб він знаходив збіги для рядків `run`, `sun`, `fun`, `pun`, `nun` і `bun`. Ваш регулярний вираз повинен використовувати символ підстановки.

# --hints--

Вам слід використати метод `.test()`.

```js
assert(__helpers.removeJSComments(code).match(/\.test\(.*\)/));
```

Вам слід використати символ підстановки у вашому регулярному виразі `unRegex`

```js
assert(/\./.test(unRegex.source));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `run` у рядку `Let us go on a run.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Let us go on a run.'));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `sun` у рядку `The sun is out today.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('The sun is out today.'));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `fun` у рядку `Coding is a lot of fun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Coding is a lot of fun.'));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `pun` у рядку `Seven days without a pun makes one weak.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('Seven days without a pun makes one weak.'));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `nun` у рядку `One takes a vow to be a nun.`

```js
unRegex.lastIndex = 0;
assert(unRegex.test('One takes a vow to be a nun.'));
```

Ваш регулярний вираз `unRegex` має знаходити збіг `bun` у рядку `She got fired from the hot dog stand for putting her hair in a bun.`

```js
unRegex.lastIndex = 0;
assert(
  unRegex.test(
    'She got fired from the hot dog stand for putting her hair in a bun.'
  )
);
```

Ваш регулярний вираз `unRegex` не має знаходити збіг у рядку `There is a bug in my code.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('There is a bug in my code.'));
```

Ваш регулярний вираз `unRegex` не має знаходити збіг у рядку `Catch me if you can.`

```js
unRegex.lastIndex = 0;
assert(!unRegex.test('Catch me if you can.'));
```

# --seed--

## --seed-contents--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /change/; // Change this line
let result = unRegex.test(exampleStr);
```

# --solutions--

```js
let exampleStr = "Let's have fun with regular expressions!";
let unRegex = /.un/; // Change this line
let result = unRegex.test(exampleStr);
```
