---
id: 587d7db6367417b2b2512b98
title: "Співпадіння одиночних символів, які не вказані"
challengeType: 1
forumTopicId: 301358
dashedName: match-single-characters-not-specified
---

# --description--

До цього моменту ви створювали множину символів, які хочете співпадати, але також можна створити множину символів, які не хочете співпадати. Такі множини символів називаються <dfn>запереченими множинами символів</dfn>.

Щоб створити заперечену множину символів, розмістіть символ карета (`^`) після відкриваючої дужки і перед символами, які не хочете співпадати.

Наприклад, `/[^aeiou]/gi` співпадає з усіма символами, які не є голосними. Зверніть увагу, що символи на кшталт `.`, `!`, `[`, `@`, `/` і порожній простір співпадають — заперечена множина голосних виключає лише голосні символи.

# --instructions--

Створіть один регулярний вираз, який співпадає з усіма символами, що не є числом або голосною. Не забудьте включити відповідні прапорці у регулярний вираз.

# --hints--

Ваш регулярний вираз `myRegex` має співпадати з 9 елементами.

```js
assert(result.length == 9);
```

Ваш регулярний вираз `myRegex` має використовувати глобальний прапорець.

```js
assert(myRegex.flags.match(/g/).length == 1);
```

Ваш регулярний вираз `myRegex` має використовувати прапорець нечутливості до регістру.

```js
assert(myRegex.flags.match(/i/).length == 1);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "3 blind mice.";
let myRegex = /change/; // Change this line
let result = myRegex; // Change this line
```

# --solutions--

```js
let quoteSample = "3 blind mice.";
let myRegex = /[^0-9aeiou]/gi; // Change this line
let result = quoteSample.match(myRegex); // Change this line
```
