---
id: 587d7db8367417b2b2512ba0
title: "Знайти все, крім літер і цифр"
challengeType: 1
forumTopicId: 301353
dashedName: match-everything-but-letters-and-numbers
---

# --description--

Ви вже знаєте, що можна використати скорочення для пошуку алфавітно-цифрових символів `[A-Za-z0-9_]` за допомогою `\w`. Природним шаблоном, який ви можете захотіти знайти, є протилежність алфавітно-цифровим символам.

Ви можете шукати протилежність `\w` за допомогою `\W`. Зверніть увагу, що протилежний шаблон використовує велику літеру. Це скорочення еквівалентне `[^A-Za-z0-9_]`.

```js
let shortHand = /\W/;
let numbers = "42%";
let sentence = "Coding!";
numbers.match(shortHand);
sentence.match(shortHand);
```

Перший виклик `match` поверне значення `["%"]`, а другий — `["!"]`.

# --instructions--

Використайте скорочений клас символів `\W`, щоб порахувати кількість неалфавітно-цифрових символів у різних цитатах і рядках.

# --hints--

Ваш регулярний вираз має використовувати глобальний прапорець.

```js
assert(nonAlphabetRegex.global);
```

Ваш регулярний вираз має знайти 6 неалфавітно-цифрових символів у рядку `The five boxing wizards jump quickly.`.

```js
assert(
  'The five boxing wizards jump quickly.'.match(nonAlphabetRegex).length == 6
);
```

Ваш регулярний вираз має використовувати скорочений клас символів для пошуку символів, які не є алфавітно-цифровими.

```js
assert(/\\W/.test(nonAlphabetRegex.source));
```

Ваш регулярний вираз має знайти 8 неалфавітно-цифрових символів у рядку `Pack my box with five dozen liquor jugs.`

```js
assert(
  'Pack my box with five dozen liquor jugs.'.match(nonAlphabetRegex).length == 8
);
```

Ваш регулярний вираз має знайти 6 неалфавітно-цифрових символів у рядку `How vexingly quick daft zebras jump!`

```js
assert(
  'How vexingly quick daft zebras jump!'.match(nonAlphabetRegex).length == 6
);
```

Ваш регулярний вираз має знайти 12 неалфавітно-цифрових символів у рядку `123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.`

```js
assert(
  '123 456 7890 ABC def GHI jkl MNO pqr STU vwx YZ.'.match(nonAlphabetRegex)
    .length == 12
);
```

# --seed--

## --seed-contents--

```js
let quoteSample = "The five boxing wizards jump quickly.";
let nonAlphabetRegex = /change/; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```

# --solutions--

```js
let quoteSample = "The five boxing wizards_jump quickly.";
let nonAlphabetRegex = /\W/g; // Change this line
let result = quoteSample.match(nonAlphabetRegex).length;
```
