---
id: 587d7db6367417b2b2512b9b
title: Знайти символи за допомогою лінивої відповідності
challengeType: 1
forumTopicId: 301341
dashedName: find-characters-with-lazy-matching
---

# --description--

У регулярних виразах <dfn>жадібне</dfn> співпадіння знаходить найдовшу можливу частину рядка, що відповідає шаблону регулярного виразу, і повертає її як співпадіння. Альтернативою є <dfn>линеве</dfn> співпадіння, яке знаходить найменшу можливу частину рядка, що задовольняє шаблон регулярного виразу.

Ви можете застосувати регулярний вираз `/t[a-z]*i/` до рядка `"titanic"`. Цей регулярний вираз фактично є шаблоном, який починається з `t`, закінчується `i` і має деякі літери між ними.

Регулярні вирази за замовчуванням жадібні, тому співпадіння поверне `["titani"]`. Воно знаходить найбільший можливий підрядок, що відповідає шаблону.

Однак ви можете використати символ `?`, щоб змінити це на линева відповідність. `"titanic"`, застосований до зміненого регулярного виразу `/t[a-z]*?i/`, повертає `["ti"]`.

**Примітка:** Парсити HTML за допомогою регулярних виразів не рекомендується, але пошук за шаблоном у рядку HTML регулярними виразами цілком прийнятний.

# --instructions--

Виправте регулярний вираз `/<.*>/`, щоб він повертав HTML-тег `<h1>`, а не текст `"<h1>Winter is coming</h1>"`. Пам’ятайте, що символ підстановки `.` у регулярному виразі відповідає будь-якому символу.

# --hints--

Змінна `result` має бути масивом, що містить `<h1>`

```js
assert(result[0] == '<h1>');
```

`myRegex` має використовувати линева відповідність

```js
assert(/[^\\][\*\+\?]\?/.test(myRegex));
```

`myRegex` не має містити рядок `h1`

```js
assert(!myRegex.source.match('h1'));
```

# --seed--

## --seed-contents--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*>/; // Change this line
let result = text.match(myRegex);
```

# --solutions--

```js
let text = "<h1>Winter is coming</h1>";
let myRegex = /<.*?>/; // Change this line
let result = text.match(myRegex);
```
