---
id: 587d7db9367417b2b2512ba5
title: Вкажіть верхню та нижню кількість співпадінь
challengeType: 1
forumTopicId: 301367
dashedName: specify-upper-and-lower-number-of-matches
---

# --description--

Нагадаємо, що ви використовуєте знак плюс `+` для пошуку одного або більше символів, а зірочку `*` — для пошуку нуля або більше символів. Це зручно, але іноді потрібно знайти співпадіння в певному діапазоні.

Ви можете вказати нижню та верхню кількість шаблонів за допомогою <dfn>кількісних специфікаторів</dfn>. Кількісні специфікатори використовуються з фігурними дужками (`{` і `}`). Між фігурними дужками ставлять два числа — для нижньої та верхньої кількості шаблонів.

Наприклад, щоб знайти лише літеру `a`, яка зустрічається від `3` до `5` разів у рядку `ah`, ваш регулярний вираз буде `/a{3,5}h/`.

```js
let A4 = "aaaah";
let A2 = "aah";
let multipleA = /a{3,5}h/;
multipleA.test(A4);
multipleA.test(A2);
```

Перший виклик `test` поверне `true`, а другий — `false`.

# --instructions--

Змініть регулярний вираз `ohRegex`, щоб він знаходив всю фразу `Oh no` лише тоді, коли в ній від `3` до `6` літер `h`.

# --hints--

Ваш регулярний вираз має використовувати фігурні дужки.

```js
assert(ohRegex.source.match(/{.*?}/).length > 0);
```

Ваш регулярний вираз не повинен знаходити рядок `Ohh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohh no'));
```

Ваш регулярний вираз має знаходити рядок `Ohhh no`

```js
assert('Ohhh no'.match(ohRegex)[0].length === 7);
```

Ваш регулярний вираз має знаходити рядок `Ohhhh no`

```js
assert('Ohhhh no'.match(ohRegex)[0].length === 8);
```

Ваш регулярний вираз має знаходити рядок `Ohhhhh no`

```js
assert('Ohhhhh no'.match(ohRegex)[0].length === 9);
```

Ваш регулярний вираз має знаходити рядок `Ohhhhhh no`

```js
assert('Ohhhhhh no'.match(ohRegex)[0].length === 10);
```

Ваш регулярний вираз не повинен знаходити рядок `Ohhhhhhh no`

```js
ohRegex.lastIndex = 0;
assert(!ohRegex.test('Ohhhhhhh no'));
```

# --seed--

## --seed-contents--

```js
let ohStr = "Ohhh no";
let ohRegex = /change/; // Change this line
let result = ohRegex.test(ohStr);
```

# --solutions--

```js
let ohStr = "Ohhh no";
let ohRegex = /Oh{3,6} no/; // Change this line
let result = ohRegex.test(ohStr);
```
