---
id: 587d7db9367417b2b2512ba7
title: Вказати точну кількість співпадінь
challengeType: 1
forumTopicId: 301365
dashedName: specify-exact-number-of-matches
---

# --description--

Ви можете вказати нижню та верхню межу кількості шаблонів за допомогою кількісних специфікаторів у фігурних дужках. Іноді потрібно, щоб співпадінь було саме певне число.

Щоб вказати точну кількість шаблонів, просто помістіть це число між фігурними дужками.

Наприклад, щоб співпадало лише слово `hah` з літерою `a` рівно `3` рази, ваш регулярний вираз буде `/ha{3}h/`.

```js
let A4 = "haaaah";
let A3 = "haaah";
let A100 = "h" + "a".repeat(100) + "h";
let multipleHA = /ha{3}h/;
multipleHA.test(A4);
multipleHA.test(A3);
multipleHA.test(A100);
```

Послідовно три виклики `test` повернули б `false`, `true` та `false`.

# --instructions--

Змініть регулярний вираз `timRegex`, щоб він співпадав зі словом `Timber` лише тоді, коли воно має чотири літери `m`.

# --hints--

Ваш регулярний вираз має використовувати фігурні дужки.

```js
assert(timRegex.source.match(/{.*?}/).length > 0);
```

Ваш регулярний вираз не повинен співпадати з рядком `Timber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timber'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Timmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmber'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Timmmber`

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Timmmber'));
```

Ваш регулярний вираз має співпадати з рядком `Timmmmber`

```js
timRegex.lastIndex = 0;
assert(timRegex.test('Timmmmber'));
```

Ваш регулярний вираз не повинен співпадати з рядком `Timber` з 30 літер `m` у ньому.

```js
timRegex.lastIndex = 0;
assert(!timRegex.test('Ti' + 'm'.repeat(30) + 'ber'));
```

# --seed--

## --seed-contents--

```js
let timStr = "Timmmmber";
let timRegex = /change/; // Change this line
let result = timRegex.test(timStr);
```

# --solutions--

```js
let timStr = "Timmmmber";
let timRegex = /Tim{4}ber/; // Change this line
let result = timRegex.test(timStr);
```
