---
id: 5c3dda8b4d8df89bea71600f
title: Перевірка змішаного групування символів
challengeType: 1
forumTopicId: 301339
dashedName: check-for-mixed-grouping-of-characters
---

# --description--

Іноді ми хочемо перевірити групи символів за допомогою регулярного виразу, і для цього ми використовуємо круглі дужки `()`.

Якщо ви хочете знайти або `Penguin`, або `Pumpkin` у рядку, ви можете використати такий регулярний вираз: `/P(engu|umpk)in/g`

Потім перевірте, чи є потрібні групи рядків у тестовому рядку, використовуючи метод `test()`.

```js
let testStr = "Pumpkin";
let testRegex = /P(engu|umpk)in/;
testRegex.test(testStr);
```

Метод `test` тут поверне `true`.

# --instructions--

Виправте регулярний вираз так, щоб він перевіряв імена `Franklin Roosevelt` або `Eleanor Roosevelt` з урахуванням регістру та допускав наявність по батькові.

Потім виправте код так, щоб створений вами регулярний вираз перевірявся на `myString`, і залежно від збігу поверталося або `true`, або `false`.

# --hints--

Ваш регулярний вираз `myRegex` має повертати `true` для рядка `Franklin D. Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Franklin D. Roosevelt'));
```

Ваш регулярний вираз `myRegex` має повертати `true` для рядка `Eleanor Roosevelt`

```js
myRegex.lastIndex = 0;
assert(myRegex.test('Eleanor Roosevelt'));
```

Ваш регулярний вираз `myRegex` має повертати `false` для рядка `Franklin Rosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Franklin Rosevelt'));
```

Ваш регулярний вираз `myRegex` має повертати `false` для рядка `Frank Roosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('Frank Roosevelt'));
```

Ваш регулярний вираз `myRegex` має повертати `false` для рядка `FranklinRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('FranklinRoosevelt'));
```

Ваш регулярний вираз `myRegex` має повертати `false` для рядка `EleanorRoosevelt`

```js
myRegex.lastIndex = 0;
assert(!myRegex.test('EleanorRoosevelt'));
```

Ви маєте використати `.test()` для тестування регулярного виразу.

```js
assert(__helpers.removeJSComments(code).match(/myRegex.test\(\s*myString\s*\)/));
```

Ваш результат має повертати `true`.

```js
assert(result === true);
```

# --seed--

## --seed-contents--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /False/; // Change this line
let result = false; // Change this line
// After passing the challenge experiment with myString and see how the grouping works
```

# --solutions--

```js
let myString = "Eleanor Roosevelt";
let myRegex = /(Franklin|Eleanor) (([A-Z]\.?|[A-Z][a-z]+) )?Roosevelt/;
let result = myRegex.test(myString);
```
