---
id: a3bfc1673c0526e06d3ac698
title: Costruisci un calcolatore della somma dei numeri primi
challengeType: 26
dashedName: build-a-prime-number-sum-calculator
---

# --description--

In questo laboratorio costruirai un calcolatore che prende un numero e restituisce la somma di tutti i numeri primi minori o uguali a quel numero.

**Obiettivo**: Completa le storie utente qui sotto e fai superare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti avere una funzione `sumPrimes` che accetta un numero come argomento.
1. La funzione `sumPrimes` dovrebbe restituire la somma di tutti i numeri primi minori o uguali al numero fornito.
1. Se il numero di input è minore di `2`, la funzione dovrebbe restituire `0`.

# --hints--

Dovresti avere una funzione `sumPrimes`.

```js
assert.isFunction(sumPrimes);
```

`sumPrimes(10)` dovrebbe restituire `17`.

```js
assert.strictEqual(sumPrimes(10), 17);
```

`sumPrimes(5)` dovrebbe restituire `10`.

```js
assert.strictEqual(sumPrimes(5), 10);
```

`sumPrimes(2)` dovrebbe restituire `2`.

```js
assert.strictEqual(sumPrimes(2), 2);
```

`sumPrimes(0)` dovrebbe restituire `0`.

```js
assert.strictEqual(sumPrimes(0), 0);
```

`sumPrimes(977)` dovrebbe restituire `73156`.

```js
assert.strictEqual(sumPrimes(977), 73156);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function sumPrimes(num) {
  function isPrime(n) {
    if (n < 2) return false;
    for (let i = 2; i <= Math.sqrt(n); i++) {
      if (n % i === 0) {
        return false;
      }
    }
    return true;
  }

  let sum = 0;

  for (let i = 2; i <= num; i++) {
    if (isPrime(i)) {
      sum += i;
    }
  }
  return sum;
}
```
