---
id: 587d824a367417b2b2512c44
title: 股票價格檢查器
challengeType: 4
forumTopicId: 301572
dashedName: stock-price-checker
---

# --description--

建置一個功能上類似於此的全端 JavaScript 應用程式：<a href="https://stock-price-checker.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker.freecodecamp.rocks/</a>。

由於所有可靠的股票價格 API 都需要 API 金鑰，我們已經建立了一個替代方案。使用 <a href="https://stock-price-checker-proxy.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://stock-price-checker-proxy.freecodecamp.rocks/</a>，即可在不需要註冊自己的金鑰的情況下取得最新的股票價格資訊。 

在這個專案中工作將涉及你使用以下其中一種方法來撰寫你的程式碼：

-   複製<a href="https://github.com/freeCodeCamp/boilerplate-project-stockchecker/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>並在區域完成你的專案。
-   使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

# --instructions--

1.  將 `NODE_ENV` 環境變數設定為 `test`，不加引號
2.  在 `routes/api.js` 中完成專案或透過建立處理常式/控制元件完成專案
3.  你將為你的 `server.js` 添加任何安全特性。
4.  你將在 `tests/2_functional-tests.js` 中創建所有的功能測試。

**注意** 隱私考量：由於要求每個 IP 只能接受 1 次按讚，你必須存檔 IP 位址。遵守像是《一般資料保護規範》這類資料隱私法規非常重要。一個選項是取得使用者資料存檔的許可，但將資料匿名化會簡單許多。針對此挑戰，請記得在將 IP 位址存檔到資料庫之前先將其匿名化。如果你需要如何做的想法，可以選擇對資料進行雜湊、截斷，或將 IP 位址的部分設為 0。

將以下測試寫入 `tests/2_functional-tests.js`：

-   查看單一股票：對 `/api/stock-prices/` 發送 GET 請求
-   查看一支股票並按讚：對 `/api/stock-prices/` 發出 GET 請求
-   再次查看相同股票並按讚：對 `/api/stock-prices/` 發出 GET 請求
-   查看兩支股票：對 `/api/stock-prices/` 的 GET 請求
-   查看兩支股票並按讚：對 `/api/stock-prices/` 發出 GET 請求

# --hints--

你可以提供你自己的專案，而非範例 URL。

```js
  assert(
    !/.*\/stock-price-checker\.freecodecamp\.rocks/.test(code)
  );
```

你應該設定內容安全政策，只允許從你的伺服器載入腳本和 CSS。

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("script-src 'self'")
  );
  assert.isTrue(
    parsed.headers['content-security-policy'].includes("style-src 'self'")
  );
```

你可以對 `/api/stock-prices` 發送 `GET` 請求，並將 NASDAQ 股票符號傳遞到 `stock` 查詢參數。傳回的物件將包含名為 `stockData` 的屬性。

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  assert.property(parsed, 'stockData');
```

`stockData` 屬性包含 `stock` 符號作為字串、`price` 作為數字，以及 `likes` 作為數字。

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker.price, 'number');
  assert.typeOf(ticker.likes, 'number');
  assert.typeOf(ticker.stock, 'string');
```

你也可以傳遞一個 `like` 欄位為 `true`（布林值），以將你的喜歡添加到股票。每個 IP 應只接受 1 個喜歡。

```js

```

如果你傳入 2 支股票，傳回的值將會是包含兩支股票資訊的陣列。它將會顯示 `rel_likes`（兩支股票的喜歡數差異），而不是 `likes`，用於兩個 `stockData` 物件。

```js
  const data = await fetch(
    code + '/api/stock-prices?stock=GOOG&stock=MSFT'
  );
  const parsed = await data.json();
  const ticker = parsed.stockData;
  assert.typeOf(ticker, 'array');
  assert.property(ticker[0], 'rel_likes');
  assert.property(ticker[1], 'rel_likes');
```

所有 5 個功能測試均已完成且通過。

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 5);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
  });
```

