---
id: 587d824a367417b2b2512c45
title: 匿名訊息板
challengeType: 4
forumTopicId: 301568
dashedName: anonymous-message-board
---

# --description--

建置一個功能上類似於此的全端 JavaScript 應用程式：<a href="https://anonymous-message-board.freecodecamp.rocks/" target="_blank" rel="noopener noreferrer nofollow">https://anonymous-message-board.freecodecamp.rocks/</a>。

在這個專案中工作將涉及你使用以下其中一種方法來撰寫你的程式碼：

-   複製<a href="https://github.com/freeCodeCamp/boilerplate-project-messageboard/" target="_blank" rel="noopener noreferrer nofollow">這個 GitHub 倉庫</a>並在區域完成你的專案。
-   使用你選擇的網站建置工具來完成專案。務必將我們 GitHub 倉庫中的所有檔案納入其中。

# --instructions--

1.  當準備好撰寫測試時，將 `NODE_ENV` 設為 test（不加引號），並將 DB 設為你的資料庫連接字串（在 `.env` 中）。
2.  建議在 `routes/api.js` 中創建控制元件／處理常式並控制路由。
3.  你將為你的 `server.js` 添加任何安全特性。

將以下測試寫入 `tests/2_functional-tests.js`：

-   建立新的執行緒：對 `/api/threads/{board}` 發出 POST 請求
-   查看每個有 3 則回覆的最新 10 個執行緒：對 `/api/threads/{board}` 的 GET 請求
-   使用錯誤密碼刪除執行緒：對 `/api/threads/{board}` 發出帶有無效 `delete_password` 的 DELETE 請求
-   使用正確密碼刪除執行緒：對 `/api/threads/{board}` 發送帶有有效 `delete_password` 的 DELETE 請求
-   回報執行緒：對 `/api/threads/{board}` 的 PUT 請求
-   建立新回覆：對 `/api/replies/{board}` 發送 POST 請求
-   查看單一執行緒及所有回覆：對 `/api/replies/{board}` 的 GET 請求
-   使用錯誤密碼刪除回覆：對 `/api/replies/{board}` 發送帶有無效 `delete_password` 的 DELETE 請求
-   使用正確密碼刪除回覆：對 `/api/replies/{board}` 發送帶有有效 `delete_password` 的 DELETE 請求
-   回報回覆：對 `/api/replies/{board}` 的 PUT 請求

# --hints--

你可以提供你自己的專案，而非範例 URL。

```js
  assert(
    !/.*\/anonymous-message-board\.freecodecamp\.rocks/.test(
      code
    )
  );
```

只允許你的網站在你自己的頁面中的 iFrame 內被載入。

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-frame-options']?.includes('SAMEORIGIN'));
```

不允許 DNS 預取。

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['x-dns-prefetch-control']?.includes('off'));
```

只允許你的网站傳送你自己頁面的 referrer。

```js
  const data = await fetch(code + '/_api/app-info');
  const parsed = await data.json();
  assert.isTrue(parsed.headers['referrer-policy']?.includes('same-origin'));
```

你可以向 `/api/threads/{board}` 發送包含 `text` 和 `delete_password` 的表單資料的 POST 請求。儲存的資料庫記錄將至少包含欄位 `_id`、`text`、`created_on`（日期與時間）、`bumped_on`（日期與時間，初始與 `created_on` 相同）、`reported`（布林值）、`delete_password` 和 `replies`（陣列）。

```js
  const date = new Date();
  const text = `fcc_test_${date}`;
  const deletePassword = 'delete_me';
  const data = { text, delete_password: deletePassword };
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(url + '/api/threads/fcc_test');
    const parsed = await checkData.json();
    try {
      assert.equal(parsed[0].text, text);
      assert.isNotNull(parsed[0]._id);
      assert.equal(new Date(parsed[0].created_on).toDateString(), date.toDateString());
      assert.equal(parsed[0].bumped_on, parsed[0].created_on);
      assert.isArray(parsed[0].replies);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以將包含 `text`、`delete_password` 和 `thread_id` 的表單資料以 POST 請求傳送到 `/api/replies/{board}`。這將會將 `bumped_on` 日期更新為留言的日期。在該討論串的 `replies` 陣列中，將會存檔一個至少具有 `_id`、`text`、`created_on`、`delete_password` 和 `reported` 屬性的物件。

```js
  const url = code;
  const body = await fetch(url + '/api/threads/fcc_test');
  const thread = await body.json();

  const date = new Date();
  const text = `fcc_test_reply_${date}`;
  const delete_password = 'delete_me';
  const thread_id = thread[0]._id;
  const replyCount = thread[0].replies.length;

  const data = { text, delete_password, thread_id };
  const res = await fetch(url + '/api/replies/fcc_test', {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  if (res.ok) {
    const checkData = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
    const parsed = await checkData.json();
    try {
      assert.equal(parsed.replies.length, replyCount + 1);
      assert.equal(parsed.replies[0].text, text);
      assert.equal(parsed._id, thread_id);
      assert.equal(parsed.bumped_on, parsed.replies[0].created_on);
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以對 `/api/threads/{board}` 發送 GET 請求。傳回將會是該看板上最近 10 個被推上的執行緒陣列，每個執行緒只包含最近 3 則回覆。`reported` 和 `delete_password` 欄位不會傳送給客戶端。

```js
  const url = code;
  const res = await fetch(url + '/api/threads/fcc_test');

  if (res.ok) {
    const threads = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isAtMost(threads.length, 10);
      for (let i = 0; i < threads.length; i++) {
        assert.containsAllKeys(threads[i], ["_id", "text", "created_on", "bumped_on", "replies"]);
        assert.isAtMost(threads[i].replies.length, 3);
        assert.notExists(threads[i].delete_password);
        assert.notExists(threads[i].reported);
        for (let j = 0; j < threads[i].replies.length; j++) {
          assert.notExists(threads[i].replies[j].delete_password);
          assert.notExists(threads[i].replies[j].reported);
        }
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以發送 GET 請求到 `/api/replies/{board}?thread_id={thread_id}`。傳回的將是包含所有回覆的整個執行緒，並且同樣排除與先前測試中客戶端相同的欄位。

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);

  if (res.ok) {
    const thread = await res.json();
    try {
      assert.equal(res.status, 200);
      assert.isObject(thread);
      assert.containsAllKeys(thread, ["_id", "text", "created_on", "bumped_on", "replies"]);
      assert.isArray(thread.replies);
      assert.notExists(thread.delete_password);
      assert.notExists(thread.reported);
      for (let i = 0; i < thread.replies.length; i++) {
        assert.notExists(thread.replies[i].delete_password);
        assert.notExists(thread.replies[i].reported);
      }
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以傳送 DELETE 請求到 `/api/threads/{board}` 並傳遞 `thread_id` 和 `delete_password` 來刪除執行緒。傳回的將會是字串 `incorrect password` 或 `success`。

```js
  const url = code;
  let res = await fetch(url + '/api/threads/fcc_test');
  const threads = await res.json();
  const thread_id = threads[0]._id;
  let data = { thread_id, delete_password: "wrong_password" };
  const res_invalid = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });
  data = { thread_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/threads/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    const not_deleted = await res_invalid.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      assert.equal(not_deleted, "incorrect password");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以傳送 DELETE 請求到 `/api/replies/{board}`，並傳遞 `thread_id`、`reply_id` 和 `delete_password`。傳回的將會是字串 `incorrect password` 或 `success`。成功時，`reply_id` 的文字將會被更改為 `[deleted]`。

```js
  const url = code;

  const thread_data = {
    text: "fcc_test_thread",
    delete_password: "delete_me",
  };
  await fetch(`${url}/api/threads/fcc_test`, {
    method: "POST",
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(thread_data)
  });
  let res = await fetch(`${url}/api/threads/fcc_test`);
  let threads = await res.json();
  const thread_id = threads[0]._id;
  
  const reply_data = { thread_id, text: "fcc_test_reply", delete_password: "delete_me" };
  await fetch(`${url}/api/replies/fcc_test`, {
    method: 'POST',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(reply_data)
  });
  res = await fetch(`${url}/api/threads/fcc_test`);
  threads = await res.json();
  const reply_id = threads[0].replies[0]._id;

  const data = { thread_id, reply_id, delete_password: "delete_me" };
  res = await fetch(url + '/api/replies/fcc_test', {
    method: 'DELETE',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const deleted = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(deleted, "success");
      res = await fetch(`${url}/api/replies/fcc_test?thread_id=${thread_id}`);
      const thread = await res.json();
      assert.equal(thread._id, thread_id);
      assert.equal(thread.replies[0]._id, reply_id);
      assert.equal(thread.replies[0].text, "[deleted]");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以傳送 PUT 請求到 `/api/threads/{board}` 並傳遞 `thread_id`。傳回的將會是字串 `reported`。`thread_id` 的 `reported` 值將會被更改為 `true`。

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const data = { thread_id };
  
  res = await fetch(`${url}/api/threads/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

你可以傳送 PUT 請求到 `/api/replies/{board}` 並傳遞 `thread_id` 和 `reply_id`。傳回的將會是字串 `reported`。`reply_id` 的 `reported` 值將會被更改為 `true`。

```js
  const url = code;

  let res = await fetch(`${url}/api/threads/fcc_test`);
  const threads = await res.json();
  const thread_id = threads[0]._id;
  const reply_id = threads[0].replies[0]._id;
  const data = { thread_id, reply_id };
  
  res = await fetch(`${url}/api/replies/fcc_test`, {
    method: 'PUT',
    headers: { 'Content-Type': 'application/json' },
    body: JSON.stringify(data)
  });

  if (res.ok) {
    const reported = await res.text();
    try {
      assert.equal(res.status, 200);
      assert.equal(reported, "reported");
    } catch (err) {
      throw new Error(err.responseText || err.message);
    }
  } else {
    throw new Error(`${res.status} ${res.statusText}`);
  }
```

所有 10 個功能測試均已完成且通過。

```js
  const tests = await fetch(code + '/_api/get-tests');
  const parsed = await tests.json();
  assert.isTrue(parsed.length >= 10);
  parsed.forEach((test) => {
    assert.equal(test.state, 'passed');
    assert.isAtLeast(test.assertions.length, 1);
  });
```

