---
id: 65f03d9f92eac9183a4d3281
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

يمكن تعريف المتجه بواسطة إحداثيين، `x` و `y`، في المستوى الإقليدي. المسافة بين نقطة الأصل لمحاور الإحداثيات والنقطة `(x, y)` ستكون طوله، أو معياره. واتجاه المتجه سيشير نحو النقطة `(x, y)`.

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

داخل صنف `Vector`، أنشئ طريقة `__init__` وأعطها ثلاث معلمات: `self` و `x` و `y`.

# --hints--

يجب تعريف طريقة `__init__` داخل صنف `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

يجب أن تأخذ طريقة `__init__` ثلاث معلمات: `self` و `x` و `y`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
