---
id: 65f058e06f34fd1f0ee6e55d
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

بعد ذلك، داخل صنف `Vector`، أنشئ طريقة `norm` فارغة وأعطها معلمة `self`.

# --hints--

يجب أن يحتوي صنف `Vector` على طريقة `norm`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

يجب أن تحتوي طريقة `norm` على معلمة `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
