---
id: 65f01969115f933073b6be03
title: خطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

المتجه هو كائن له طول (أو مقدار) واتجاه، ولا يمكن التعبير عنه برقم واحد فقط. في الفيزياء، تُستخدم المتجهات عادةً لتمثيل القوى، والسرعات، والتسارعات، وكميات أخرى.

في هذا المشروع ستبني فضاء متجهات، وهو مجموعة تُعرّف فيها سلسلة من العمليات بين العناصر داخل تلك المجموعة. ستتعلم كل التفاصيل قريبًا جدًا.

ابدأ المشروع الآن بإعلان صنف فارغ اسمه `Vector`.

# --hints--

يجب إعلان صنف فارغ اسمه `Vector`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
