---
id: 65f055b9190fc41ca35549b8
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

تقدم Python طرقًا مختلفة تتضمن شرطة سفلية مزدوجة في بداية ونهاية أسمائها. قد تكون مألوفًا بالفعل مع بعضها، مثل `__init__` و`__str__`. تُسمى هذه الطرق، التي تتبع نمط التسمية `__<name>__`، الطرق الخاصة أو الطرق السحرية أو طرق *dunder* (اختصار double underscore).

تعريف الطرق الخاصة في صنف يؤثر على سلوك ذلك الصنف. تُستدعى هذه الطرق ضمنيًا في مواقف محددة (مثل `__init__` أثناء إنشاء الكائن، و`__str__` عند طباعة الكائن أو تمريره إلى `str()`). في هذا المشروع، ستتعلم بعضًا من أكثر الطرق الخاصة استخدامًا.

في الوقت الحالي، عيّن `x` إلى الخاصية `x` لكائن `Vector`.

# --hints--

يجب تعيين `x` إلى `self.x`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
