---
id: 65f0694fb2296f3caadf8347
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

طول المتجه $\mathbf{a}$، أو المعيار، يُشار إليه عادةً بـ $\\| \mathbf{a} \\|$. يمكن حسابه كـ الجذر التربيعي لمجموع مربعات مكوناته:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

احسب معيار المتجه وأرجع النتيجة من `norm` الطريقة الخاصة بك.

# --hints--

يجب أن تستخدم `norm` الطريقة الصيغة المقدمة لحساب وإرجاع رقم يمثل معيار نسخة المتجه الحالية. لا تقرب القيمة.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
