---
id: 65f3fa097f26b510db6c710b
title: الخطوة 9
challengeType: 20
dashedName: step-9
---

# --description--

رائع، `norm` يعمل كما هو متوقع. الآن، إذا حاولت طباعة `v1`، ستحصل على التمثيل النصي الافتراضي للكائن (شيء مثل `<__main__.Vector object at 0x11eb778>`).

داخل صنف `Vector`، أعلن عن طريقة `__str__` فارغة لتنفيذ تمثيل نصي قابل للقراءة. تذكر أن تعطيها معلمة `self`.

انتبه ألا تطبع `v1` قبل أن تُرجع `__str__` سلسلة نصية، وإلا ستحصل على `TypeError`، لأن `__str__` يجب أن تُرجع دائمًا سلسلة نصية.

# --hints--

يجب أن تعرف طريقة `__str__` داخل صنف `Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

يجب أن تحتوي طريقة `__str__` على معلمة `self`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
