---
id: 65f3fe07cc763212efe91285
title: الخطوة 11
challengeType: 20
dashedName: step-11
---

# --description--

يمكن أن يحتوي المتجه على عدد `n` من الأبعاد (المكوّنات). إليك تمثيل لمتجه ثلاثي الأبعاد:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

حتى الآن، أنشأت متجهًا ثنائي الأبعاد. تريد أن تتمكن من تمثيل المتجهات بعدد مختلف من الأبعاد دون إعادة كتابة الشفرة اللازمة لكل حالة محددة. لهذا، ستستخدم الوراثة.

ابدأ بإعادة تسمية صنف `Vector` إلى `R2Vector` لتحديد أن هذا الصنف سيتعامل مع المتجهات ثنائية الأبعاد. تذكر تعديل إنشاء الكائن `v1` أيضًا.

# --hints--

يجب إعادة تسمية صنف `Vector` إلى `R2Vector`. تذكر تعديل إنشاء الكائن `v1` أيضًا.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
