---
id: 65f3fa60a93b84110b3f1708
title: الخطوة 10
challengeType: 20
dashedName: step-10
---

# --description--

تريد إرجاع تمثيل نصي ذي معنى. على سبيل المثال، في حالة `v1`، تريد إرجاع سلسلة نصية تحتوي على الصف المرتب بقيم مكوِّنات المتجه: `(2, 3)`.

من الطريقة `__str__`، أرجع سلسلة نصية تمثل المتجه كصف مرتب يحتوي مكوِّنات المتجه بالترتيب. ثم، اخرج من صنف `Vector` واطبع `v1` للتحقق من النتيجة.

# --hints--

يجب أن تُرجع الطريقة `__str__` سلسلة نصية تمثل المتجه كصف مرتب يحتوي مكوِّنات المتجه بالترتيب.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

يجب أن تطبع `v1`.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
