---
id: 65f4520e363e2642f8112e33
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

كما ترى، هناك شيء غير صحيح تمامًا. الطريقتان `norm` و `__str__` الموروثتان من `R2Vector` لا يمكنهما التكيف مع متجه ثلاثي الأبعاد. يجب أن يكون تنفيذهما أكثر مرونة.

كل كائن في Python يحتوي على خاصية خاصة تُسمى `__dict__`، وهي قاموس يخزن خصائص الكائن.

احذف استدعاءات `print` الحالية. ثم اطبع الخاصية `__dict__` لمتجهات `v1` و `v2` لترى شكلها.

# --hints--

يجب ألا يحتوي شفرتك على `print(v2)` و `print(v2.norm())`.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

يجب أن تطبع الخاصية `__dict__` لكل من `v1` و `v2`.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
