---
id: 65f407ea37ad6e181b90462e
title: الخطوة 14
challengeType: 20
dashedName: step-14
---

# --description--

يمكنك تعيين كل معلمة `i` إلى `self.i` كما فعلت سابقًا. رغم أن عدد الأسطر التي تحتاج لتكرارها هنا قليل، فإن إحدى الطرق لتجنب هذا التكرار هي استخدام الدالة `super()`. تتيح لك `super()` الإشارة ضمنيًا إلى الصنف الأب: `super().__init__(x, y)` تستدعي طريقة `__init__` الخاصة بالصنف الأب.

أضف استدعاء `super().__init__(x, y)` إلى طريقة `__init__` الخاصة بك.

# --hints--

يجب أن يحتوي طريقة `__init__` الخاصة بك على استدعاء `super().__init__(x, y)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
