---
id: 65f40051d6b09a139f253e8e
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

تتيح الوراثة تعريف صنف جديد من صنف موجود. الصنف الجديد، المسمى الطفل، يرث كل الطرق والخصائص من الصنف الموجود، المسمى الأصل.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

في المثال أعلاه، يرث الصنف الطفل `Oak` من `Tree` ويرث الطريقة `sprout` من الصنف الأصل `Tree`.

أنشئ صنفًا جديدًا باسم `R3Vector` واتبع المثال أعلاه لجعله يرث من الصنف `R2Vector`.

# --hints--

يجب تعريف صنف جديد باسم `R3Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

يجب أن يرث الصنف الجديد `R3Vector` من `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
