---
id: 65f412c208c3791fee305acf
title: الخطوة 17
challengeType: 20
dashedName: step-17
---

# --description--

في Python، يمكنك فرض استخدام المعلمات التي تُمرر بالكلمة المفتاحية فقط عن طريق إضافة `*` كمعلمة إضافية في توقيع الدالة أو الطريقة. عدل كلا الطريقتين `__init__` بإضافة `*` كالمعلمة الثانية (بعد `self`). كل معلمة توضع بعد ذلك ستتطلب استخدام معلمة بالكلمة المفتاحية عند استدعاء الدالة/الطريقة.

هذا يعني أنك تحتاج أيضًا إلى تعديل استدعاء `super().__init__(x, y)`. افعل ذلك بإعطاء `x` القيمة `x`، و`y` القيمة `y`.

أخيرًا، عدل إنشاء `v1` و`v2` باستخدام معلمات بالكلمة المفتاحية.

# --hints--

يجب فرض استخدام المعلمات بالكلمة المفتاحية عن طريق إضافة `*` كالمعلمة الثانية في كلا الطريقتين `__init__`.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

يجب تعديل استدعاء `super.__init__(x, y)` لاستخدام معلمات بالكلمة المفتاحية.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

يجب تعديل تعيين `v1` ليكون `R2Vector(x=2, y=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

يجب تعديل تعيين `v2` ليكون `R3Vector(x=2, y=2, z=3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
