---
id: 65f40401e6ef53173c04e27d
title: الخطوة 13
challengeType: 20
dashedName: step-13
---

# --description--

داخل الصنف الجديد، أعلن عن طريقة `__init__`. اجعل لها المعلمات `self` و `x` و `y` و `z`.

# --hints--

يجب أن تعرف طريقة `__init__` داخل صنف `R3Vector` الخاص بك.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

يجب أن تحتوي طريقة `__init__` على أربع معلمات هي `self` و `x` و `y` و `z`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
