---
id: 65f40fdb3579aa1ced28b2eb
title: الخطوة 16
challengeType: 20
dashedName: step-16
---

# --description--

أنشئ مثيلًا من `R3Vector` بالقيم `2` و `2` و `3` لـ `x` و `y` و `z` على التوالي. عيّن المثيل الجديد إلى `v2`.

# --hints--

يجب إعلان متغير `v2` وتعيينه إلى `R3Vector(2, 2, 3)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(2, 2, 3)")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--

--fcc-editable-region--
```
