---
id: 65f876d17832001e8e1abb05
title: الخطوة 28
challengeType: 20
dashedName: step-28
---

# --description--

أعلن عن متغير `arg_list`. عيّن له قيمة تعبير قائمة يقوم بالتكرار على `key` و `val` لكل زوج مفتاح-قيمة في `vars(self).items()` ويحسب سلسلة f-string `f'{key}={val}'` في كل تكرار.

# --hints--

يجب إعلان متغير `arg_list`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

يجب تعيين تعبير قائمة يقوم بالتكرار على `vars(self).items()` إلى `arg_list`.

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

يجب أن يستخدم تعبير القائمة المعين إلى `arg_list` المتغيرين `key` و `val` للتكرار على `vars(self).items()`.

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

يجب أن يحسب تعبير القائمة المعين إلى `arg_list` السلسلة `f'{key}={val}'` لكل زوج مفتاح-قيمة في `vars(self).items()`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
