---
id: 65f8709620b2ce1a62608f5a
title: الخطوة 23
challengeType: 20
dashedName: step-23
---

# --description--

عندما تحتاج إلى وصول ديناميكي لبعض الخصائص بدءًا من مُدخَل سلسلة نصية، فإن الدالة المدمجة `getattr()` هي ما تحتاجه. تأخذ كائنًا كمعلمتها الأولى، وسلسلة نصية تحتوي على اسم الخاصية كمعلمتها الثانية.

ابدأ بإصلاح طريقة `__str__` عن طريق استبدال السلسلة النصية التي تُرجعها `__str__()` بتعبير مولِّد يكرر عبر خصائص الكائن ويستدعي `getattr()` لكل خاصية `i`.

# --hints--

يجب أن تُرجع تعبيرًا مولِّدًا يكرر عبر `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

يجب أن تُرجع تعبيرًا مولِّدًا يستخدم `i` كمتغير التكرار.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

يجب أن تُرجع تعبيرًا مولِّدًا يستدعي `getattr(self, i)` لكل عنصر `i` في `vars(self)`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
