---
id: 65f9bc10a9fb1612e066e9e8
title: الخطوة 42
challengeType: 20
dashedName: step-42
---

# --description--

عند جمع متجهين، يُضاف كل مكوِّن من أحد المتجهين إلى نفس المكوِّن في المتجه الآخر. على سبيل المثال، جمع `(1, 2)` و `(2, 4)` يولد متجهًا ثالثًا `(3, 6)`، حيث `3` هو مجموع `1` و `2` و `6` هو مجموع `2` و `4`.

ستحتوي قاموس `kwargs` على أزواج المفتاح-القيمة اللازمة لإنشاء متجه جديد من نفس صنف المتجهين المجمَّعين.

حوّل القاموس الفارغ إلى تعبير قاموس يقوم بالتكرار على `vars(self)` ولكل مفتاح (`i`) ينشئ زوج مفتاح-قيمة، حيث يكون المفتاح هو نفس مفتاح التكرار الحالي وقيمته هي مجموع قيم خاصيتي `i` للمعاملين.

# --hints--

يجب أن تحوّل القاموس الفارغ المُعيَّن إلى `kwargs` إلى تعبير قاموس يتكرر على `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

يجب أن يستخدم تعبير القاموس المُعيَّن إلى `kwargs` المتغير `i` للتكرار على `vars(self)`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

يجب أن يستخدم تعبير القاموس المُعيَّن إلى `kwargs` المفتاح `i`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

يجب أن يُعيّن تعبير القاموس المُعيَّن إلى `kwargs` مجموع خاصية `i` في `self` وخاصية `i` في `other` إلى المفتاح `i` لكل `i` في `vars(self)`. استخدم الدالة `getattr()` لذلك.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
