---
id: 65f9dd6e5a08af19c196c2df
title: الخطوة 45
challengeType: 20
dashedName: step-45
---

# --description--

يُحصل على المتجه الناتج من طرح متجه من آخر بحساب الفرق بين كل مكوِّن منهما. على سبيل المثال، طرح `(2, 4)` من `(7, 3)` يولد متجهًا ثالثًا `(5, -1)`، حيث `5` هو الفرق بين `7` و `2` و `-1` هو الفرق بين `3` و `4`.

كما فعلت سابقًا داخل طريقة `__add__`، أعلن متغير `kwargs` بعد جملة `if`. عيّنه إلى تعبير قاموسي يتكرّر على خصائص الكائن ويُنشئ زوج مفتاح-قيمة لكل مفتاح `i`، حيث يكون المفتاح هو نفس المفتاح في التكرار الحالي وقيمته هي الفرق بين قيم خاصيتي `i` للمعاملين.

# --hints--

يجب إعلان متغير `kwargs` وتعيينه إلى تعبير قاموسي يتكرّر على خصائص الكائن.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

يجب أن يستخدم التعبير القاموسي المعين إلى `kwargs` المتغير `i` للتكرار على خصائص الكائن.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

يجب أن يستخدم التعبير القاموسي المعين إلى `kwargs` المفتاح `i` كمفتاح.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

يجب أن يعين التعبير القاموسي المعين إلى `kwargs` الفرق بين خاصية `i` في `self` وخاصية `i` في `other` إلى المفتاح `i` لكل `i`.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
