---
id: 65f9cb678070ca1668898c70
title: الخطوة 44
challengeType: 20
dashedName: step-44
---

# --description--

بنفس الطريقة التي يُستدعى فيها `__add__` ضمنيًا عند جمع كائنين معًا، تُستدعى طريقة `__sub__` ضمنيًا في حالة الطرح.

الآن، عرّف طريقة `__sub__` فارغة وأعطها معلمتين: `self` و`other`. داخل طريقتك الجديدة، أنشئ جملة if الشرطية للتحقق مما إذا كان `self` و`other` لا ينتميان إلى نفس الصنف وأرجع `NotImplemented`، كما فعلت سابقًا.

# --hints--

يجب أن تعرّف طريقة `__sub__` داخل صنف `R2Vector`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

يجب أن تحتوي طريقة `__sub__` على معلمتين، `self` و`other`.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

يجب أن تنشئ جملة `if` تتحقق مما إذا كان `self` و`other` لا ينتميان إلى نفس الصنف.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

يجب أن تُعيد `NotImplemented` من داخل جسم جملة `if`.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
