---
id: 65faaed8f7a9772f023ea816
title: الخطوة 52
challengeType: 20
dashedName: step-52
---

# --description--

يمكن ضرب المتجهات في *عدد قياسي*، أي رقم يضرب كل مكوِّن منفرد. ناتج الضرب في العدد القياسي هو متجه له نفس اتجاه المتجه الأصلي لكن بقيمة مقدار مختلفة.

نفذ الضرب في العدد القياسي بالتحقق مما إذا كان `other` إما `int` أو `float`. إذا كان كذلك، أرجع نسخة جديدة من الصنف الحالي تحتوي على كل مكوِّن من المتجه الأصلي مضروبًا في `other`. هذا سيكون المتجه الناتج من الضرب في العدد القياسي.

تأكد من أن الطرق يمكن تطبيقها لحساب الضرب في العدد القياسي للمتجهات ذات أي عدد من الأبعاد.

# --hints--

يجب أن تُرجع طريقتك نسخة جديدة من الصنف الحالي فقط عندما يكون نوع `other` إما `int` أو `float`.

```js
({ test: () => runPython(`
v1 = R2Vector(x=0, y=0)
v2 = R3Vector(x=0, y=0, z=0)
types = ["", True, [], {}]
assert all((v1 * i) is None for i in types)
types = [1, 1.0]
assert all(type(v1 * i) is type(v1) for i in types)
assert all(type(v2 * i) is type(v2) for i in types)
`) })
```

يجب أن يحتوي المتجه الناتج من الضرب في العدد القياسي على كل مكوِّن من المتجه الأصلي مضروبًا في العدد القياسي.

```js
({ test: () => runPython(`
v1 = R2Vector(x=1, y=1.5)
v2 = R3Vector(x=2.2, y=3, z=-1)
assert vars(v1 * 3) == vars(R2Vector(x=3, y=4.5))
assert vars(v2 * (-2.0)) == vars(R3Vector(x=-4.4, y=-6.0, z=2.0))
`) })
```


# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __mul__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
