---
id: 666832e427d70bc5219dc62a
title: الخطوة 36
challengeType: 20
dashedName: step-36
---

# --description--

كما ترى من المخرجات، رغم أنك عرّفت `__getattr__` في الصنف الخاص بك، هذه الطريقة لم تُستدعَ بعد. يحدث هذا لأن الوصول الافتراضي إلى الخاصية يتم عبر `__getattribute__`. لذلك، ترى قيمة الخاصية مطبوعة على الطرفية.

الآن عدّل السطرين الأخيرين من الشفرة للوصول إلى الخاصية `z` في `v1` في كلا استدعائي `print()`. هذه المرة، ستصل إلى خاصية لا يمتلكها `v1`. نتيجة لذلك، سترفع `__getattribute__` خطأً وسيتم استدعاء `__getattr__`.

# --hints--

يجب تعديل الشفرة لطباعة الخاصية `z` في `v1` باستخدام كل من عامل النقطة والدالة `getattr()`.

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
