---
id: 669e2f60e83c011754f711f9
title: بناء صفحة وكالة سفر
challengeType: 25
dashedName: build-a-travel-agency-page
demoType: onClick
---

# --description--

**الهدف:** إكمال قصص المستخدم أدناه واجتياز جميع الاختبارات لإتمام المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي على إعلان `DOCTYPE`.
1. يجب أن تحتوي على عنصر `html` مع تعيين `lang` إلى `en`.
1. يجب أن تحتوي على عنصر `head` يحتوي على عنصر `meta` فارغ مع تعيين `charset` إلى `utf-8` وعنصر `title` بنص `Travel Agency Page`.
1. يجب أن يحتوي عنصر `head` على وسم `meta` يحتوي على وصف مختصر لموقعك لأغراض تحسين محركات البحث.
1. يجب أن يحتوي على عنصر `h1` لعرض وجهات السفر الخاصة بك. يجب أن يكون العنصر `h1` الوحيد في الصفحة.
1. يجب أن يحتوي على فقرة أسفل عنصر `h1` لتعريف فرص السفر.
1. يجب أن يحتوي على عنصر `h2` بنص `Packages`.
1. يجب أن يحتوي على عنصر `p` يقدم باختصار الحزم المختلفة.
1. يجب أن يحتوي على عنصر قائمة غير مرتبة به عنصران في القائمة. يجب أن يحتوي العنصران على نصي `Group Travels` و`Private Tours` على التوالي. يجب أن يكون نص كل عنصر قائمة محاطًا بعنصر وسم الرابط.
1. يجب أن يحتوي على عنصر `h2` بنص `Top Itineraries`.
1. يجب أن يحتوي على ثلاثة عناصر `figure` على الأقل، يحتوي كل منها على عنصر وسم الرابط وعنصر `figcaption` يصف صورة الرسم البياني.
1. يجب أن تحتوي عناصر وسم الرابط الثلاثة على عنصر `img` مع خاصية `alt` مناسبة وخاصية `src` تشير إلى صورة صالحة كمحتوى لها. يمكنك استخدام `https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg`، `https://cdn.freecodecamp.org/curriculum/labs/alps.jpg`، و`https://cdn.freecodecamp.org/curriculum/labs/sea.jpg` إذا أردت.
1. يجب أن تحتوي جميع عناصر وسم الرابط الخمسة على خاصية `href` بقيمة `https://www.freecodecamp.org/learn` وخاصية `target` بقيمة `_blank`.

# --hints--

يجب أن تحتوي صفحة وكالة السفر على إعلان `!DOCTYPE html`.

```js
assert.match(code, /<!DOCTYPE\s+html>/i);
```

يجب أن يكون لديك عنصر `html` مع تعيين `lang` إلى `en`.

```js
assert.match(code, /<\s*html(?=[\s>])[^>]*\slang\s*=\s*('|")en\1[^>]*>/i);
```

يجب أن يحتوي عنصر `html` على وسم إغلاق.

```js
assert.match(code, /<\/html\s*>/i);
```

يجب أن يكون لديك عنصر `head` داخل عنصر `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>/i
);
```

يجب أن يحتوي عنصر `head` على وسم الإغلاق.

```js
assert.match(code, /<\/head\s*>/i);
```

يجب أن يحتوي عنصر `head` على عنصر `title`.

```js
assert.isNotNull(document.querySelector('head title'));
```

يجب أن يحتوي عنصر `title` على وسم الإغلاق.

```js
assert.match(code, /<\/title\s*>/i);
```

يجب أن يكون لديك عنصران `meta` داخل عنصر `head`.

```js
assert.isAtLeast(document.querySelectorAll('head meta').length, 2);
```

يجب أن يحتوي أحد عناصر `meta` على خاصية `name` بقيمة `description` وخاصية `content` غير فارغة.

```js
assert(
  Array.from(document.querySelectorAll('head meta')).some(
    (meta) =>
      meta.getAttribute('name')?.toLowerCase() === 'description' &&
      meta.getAttribute('content')?.trim()
  )
);
```

يجب أن يحتوي أحد عناصر `meta` على خاصية `charset` بقيمة `utf-8`.

```js
const charsetMeta = document.querySelector('head meta[charset]');
assert.equal(charsetMeta?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

يجب أن يحتوي عنصر `title` على اسم وكالة السفر الخاصة بك.

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

يجب أن يكون لديك عنصر `body` داخل عنصر `html`.

```js
assert.match(
  code,
  /<\s*html(?=[\s>])[^>]*>(?:(?!<\/html\s*>)[\s\S])*<\s*head\s*>(?:(?!<\/html\s*>)[\s\S])*<\s*body\s*>/i
);
```

يجب أن يحتوي عنصر `body` على وسم الإغلاق.

```js
assert.match(code, /<\/body\s*>/i);
```

يجب أن يكون لديك عنصر `h1` لعرض وجهات السفر الخاصة بك.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

يجب أن يكون لديك عنصر `h1` واحد فقط.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

يجب أن يكون لديك عنصر `p` أسفل عنصر `h1`.

```js
assert.strictEqual(
  document.querySelector('h1')?.nextElementSibling.tagName, 'P'
);
```

يجب أن يقدم أول عنصر `p` فرص السفر.

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

يجب أن يحتوي أول عنصر `h2` على النص `Packages`.

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText.trim(), 'Packages');
```

يجب أن يكون لديك عنصر `p` أسفل أول عنصر `h2`.

```js
assert.strictEqual(
    document.querySelector('h2')?.nextElementSibling.tagName, 'P'
);
```

يجب أن يقدم ثاني عنصر `p` باختصار الحزم المختلفة.

```js
assert.isNotEmpty(document.querySelectorAll('p')?.[1]?.textContent?.trim());
```

يجب أن يكون لديك عنصر قائمة غير مرتبة أسفل ثاني عنصر `p`.

```js
assert.strictEqual(
  document.querySelectorAll('p')?.[1]?.nextElementSibling?.tagName,
  'UL'
);
```

يجب أن يحتوي عنصر القائمة غير المرتبة على عنصرين.

```js
assert.lengthOf(document.querySelectorAll('p + ul li'), 2);
```

يجب أن يحتوي كلا العنصرين في القائمة على عنصر وسم الرابط.

```js
const listItems = document.querySelectorAll('p + ul li');
assert.isNotEmpty(listItems);
for (let e of listItems) {
  assert.strictEqual(e.children[0].tagName, 'A');
  assert.lengthOf(e.children, 1);
}
```

يجب أن يحيط عنصر وسم الرابط في أول عنصر قائمة بالنص `Group Travels`.

```js
assert.equal(document.querySelectorAll('li > a')[0]?.innerText.trim(), 'Group Travels');
```

يجب أن يحيط عنصر وسم الرابط في ثاني عنصر قائمة بالنص `Private Tours`.

```js
assert.equal(document.querySelectorAll('li > a')[1]?.innerText.trim(), 'Private Tours');
```

يجب أن يكون لديك عنصر `h2` بعد القائمة غير المرتبة.

```js
assert.strictEqual(
  document.querySelectorAll('h2')[1]?.previousElementSibling.tagName,
  'UL'
);
```

يجب أن يحتوي ثاني عنصر `h2` على النص `Top Itineraries`.

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText.trim(), 'Top Itineraries');
```

يجب أن يكون لديك ثلاثة عناصر `figure` على الأقل.

```js
assert.isAtLeast(document.querySelectorAll('figure').length, 3);
```

يجب أن يحتوي كل عنصر `figure` على عنصر وسم الرابط كأول عنصر مسكن.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[0].tagName, 'A');
}
```

يجب أن يحتوي كل عنصر `figure` على عنصر `figcaption` كثاني عنصر مسكن.

```js
const figures = document.querySelectorAll('figure');
assert.isNotEmpty(figures);
for (let figure of figures) {
  assert.equal(figure.children[1]?.tagName, 'FIGCAPTION');
}
```

يجب أن يحتوي كل عنصر `figcaption` على بعض النص.

```js
const figcaptionEls = document.querySelectorAll('figcaption');
assert.isNotEmpty(figcaptionEls);
figcaptionEls.forEach((figcaption => assert.isNotEmpty(figcaption.innerText)));
```

يجب أن يحتوي كل عنصر `a` المسكن داخل عناصر `figure` على صورة.

```js
const anchors = document.querySelectorAll('figure > a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.children[0]?.tagName, 'IMG');
}
```

يجب أن يحتوي كل عنصر `img` على خاصية `src` صالحة.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.src.trim().length, 0);
}
```

يجب أن يحتوي كل عنصر `img` على خاصية `alt` بقيمة مناسبة.

```js
const images = document.querySelectorAll('img');
assert.isNotEmpty(images);
for (let image of images) {
  assert.isAbove(image.alt.trim().length, 0);
}
```

يجب أن يحتوي كل عنصر `a` على خاصية `href` بقيمة `https://www.freecodecamp.org/learn`. لا تنسَ الروابط في عناصر القائمة.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.href, 'https://www.freecodecamp.org/learn');
}
```

يجب أن يحتوي كل عنصر `a` على خاصية `target` بقيمة `_blank`. لا تنسَ الروابط في عناصر القائمة.

```js
const anchors = document.querySelectorAll('a');
assert.isNotEmpty(anchors);
for (let anchor of anchors) {
  assert.equal(anchor.target, '_blank');
}
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Travel Agency Page</title>
    <meta name="description" content="This is an example website."/>
  </head>
  <body>
    <h1>Discover Italy</h1>
    <p>
      Art, folklore, food, nature, and more. Choose among our wide selection of
      guided tours and excursions, and live an unforgettable experience
      exploring Italy.
    </p>
    <h2>Packages</h2>
    <p>
      We offer an extensive range of holiday solutions to accommodate the needs
      of all our clients. From daily excursions in the most beautiful cities, to
      thorough tours of hidden villages and medieval towns to discover Italy's
      lesser-known sides.
    </p>
    <ul>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Group Travels</a
        >
      </li>
      <li>
        <a href="https://www.freecodecamp.org/learn" target="_blank"
          >Private Tours</a
        >
      </li>
    </ul>
    <h2>Top Itineraries</h2>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/colosseo.jpg"
          alt="colosseum view"
        />
      </a>
      <figcaption>Rome and Center Italy</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/alps.jpg"
          alt="dolomites mountain view"
        />
      </a>
      <figcaption>Nature and National Parks</figcaption>
    </figure>
    <figure>
      <a href="https://www.freecodecamp.org/learn" target="_blank">
        <img
          src="https://cdn.freecodecamp.org/curriculum/labs/sea.jpg"
          alt="sea view"
        />
      </a>
      <figcaption>South Italy and Islands</figcaption>
    </figure>
  </body>
</html>
```
