---
id: 587d7db0367417b2b2512b83
title: استخدم الوراثة حتى لا تكرر نفسك
challengeType: 1
forumTopicId: 301334
dashedName: use-inheritance-so-you-dont-repeat-yourself
---

# --description--

هناك مبدأ في البرمجة يُسمى <dfn>لا تكرر نفسك (DRY)</dfn>. سبب كون الشفرة المكررة مشكلة هو أن أي تغيير يتطلب تعديل الشفرة في أماكن متعددة. هذا يعني عادةً المزيد من العمل للمبرمجين وزيادة احتمال حدوث أخطاء.

لاحظ في المثال أدناه أن الطريقة `describe` مشتركة بين `Bird` و `Dog`:

```js
Bird.prototype = {
  constructor: Bird,
  describe: function() {
    console.log("My name is " + this.name);
  }
};

Dog.prototype = {
  constructor: Dog,
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

طريقة `describe` مكررة في مكانين. يمكن تعديل الشفرة لتتبع مبدأ DRY بإنشاء `supertype` (أو أصل) يسمى `Animal`:

```js
function Animal() { };

Animal.prototype = {
  constructor: Animal, 
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

بما أن `Animal` يحتوي على طريقة `describe`، يمكنك إزالتها من `Bird` و `Dog`:

```js
Bird.prototype = {
  constructor: Bird
};

Dog.prototype = {
  constructor: Dog
};
```

# --instructions--

طريقة `eat` مكررة في كل من `Cat` و `Bear`. عدل الشفرة بروح مبدأ DRY بنقل طريقة `eat` إلى `supertype` الخاص بـ `Animal`.

# --hints--

يجب أن يحتوي `Animal.prototype` على الخاصية `eat`.

```js
assert(Animal.prototype.hasOwnProperty('eat'));
```

يجب ألا يحتوي `Bear.prototype` على الخاصية `eat`.

```js
assert(!Bear.prototype.hasOwnProperty('eat'));
```

يجب ألا يحتوي `Cat.prototype` على الخاصية `eat`.

```js
assert(!Cat.prototype.hasOwnProperty('eat'));
```

# --seed--

## --seed-contents--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,

};
```

# --solutions--

```js
function Cat(name) {
  this.name = name;
}

Cat.prototype = {
  constructor: Cat
};

function Bear(name) {
  this.name = name;
}

Bear.prototype = {
  constructor: Bear
};

function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
```
