---
id: 587d7db1367417b2b2512b87
title: إضافة طرق بعد الوراثة
challengeType: 1
forumTopicId: 301315
dashedName: add-methods-after-inheritance
---

# --description--

يمكن لدالة البناء التي ترث كائن `prototype` الخاص بها من دالة بناء من النوع الأعلى أن تمتلك طرقها الخاصة بالإضافة إلى الطرق الموروثة.

على سبيل المثال، `Bird` هي دالة بناء ترث `prototype` الخاص بها من `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
Bird.prototype.constructor = Bird;
```

بالإضافة إلى ما يُورث من `Animal`، قد ترغب في إضافة سلوك خاص لكائنات `Bird`. هنا، ستحصل `Bird` على دالة `fly()`. تُضاف الدوال إلى `Bird's` `prototype` بنفس طريقة أي دالة بناء:

```js
Bird.prototype.fly = function() {
  console.log("I'm flying!");
};
```

الآن سيكون لدى مثيلات `Bird` كل من الطريقتين `eat()` و `fly()`:

```js
let duck = new Bird();
duck.eat();
duck.fly();
```

ستعرض `duck.eat()` السلسلة النصية `nom nom nom` في وحدة التحكم، وستعرض `duck.fly()` السلسلة النصية `I'm flying!`.

# --instructions--

أضف كل الشفرة اللازمة لكي يرث كائن `Dog` من `Animal` ويُضبط دالة البناء في `prototype` الخاص بـ `Dog` على `Dog`. ثم أضف طريقة `bark()` إلى كائن `Dog` بحيث يستطيع `beagle` أن يستخدم كل من `eat()` و `bark()`. يجب أن تطبع طريقة `bark()` النص `Woof!` في وحدة التحكم.

# --hints--

يجب ألا يستجيب `Animal` للطريقة `bark()`.

```js
assert(typeof Animal.prototype.bark == 'undefined');
```

يجب أن يرث `Dog` طريقة `eat()` من `Animal`.

```js
assert(typeof Dog.prototype.eat == 'function');
```

يجب أن يحتوي النموذج الأولي لـ `Dog` على طريقة `bark()`.

```js
assert('bark' in Dog.prototype);
```

يجب أن يكون `beagle` من نوع `instanceof` لـ `Animal`.

```js
assert(beagle instanceof Animal);
```

يجب أن تُضبط دالة البناء لـ `beagle` على `Dog`.

```js
assert(beagle.constructor === Dog);
```

يجب أن تسجل `beagle.eat()` السلسلة النصية `nom nom nom`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.eat();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'nom nom nom');
```

يجب أن تسجل `beagle.bark()` السلسلة النصية `Woof!`

```js
const spy = __helpers.spyOn(console, 'log');
beagle.bark();
assert.lengthOf(spy.calls,1);
assert(spy.calls[0][0] == 'Woof!');
```

# --seed--

## --seed-contents--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }

// Only change code below this line




// Only change code above this line

let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
Animal.prototype.eat = function() { console.log("nom nom nom"); };

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Dog.prototype.bark = function () {
  console.log('Woof!');
};
let beagle = new Dog();

beagle.eat();
beagle.bark();
```
