---
id: 587d7dad367417b2b2512b78
title: استخدم دالة البناء لإنشاء كائنات
challengeType: 1
forumTopicId: 18233
dashedName: use-a-constructor-to-create-objects
---

# --description--

إليك دالة البناء `Bird` من التحدٍّ السابق:

```js
function Bird() {
  this.name = "Albert";
  this.color  = "blue";
  this.numLegs = 2;
}

let blueBird = new Bird();
```

**ملاحظة:** تشير `this` داخل دالة البناء دائمًا إلى الكائن الجاري إنشاؤه.

لاحظ أن عامل `new` يُستخدم عند استدعاء دالة البناء. هذا يخبر JavaScript بإنشاء نسخة جديدة من `Bird` تُسمى `blueBird`. بدون عامل `new`، لن تشير `this` داخل دالة البناء إلى الكائن الجديد، مما يؤدي إلى نتائج غير متوقعة. الآن يحتوي `blueBird` على كل الخصائص المعرفة داخل دالة البناء `Bird`:

```js
blueBird.name;
blueBird.color;
blueBird.numLegs;
```

مثل أي كائن آخر، يمكن الوصول إلى خصائصه وتعديلها:

```js
blueBird.name = 'Elvira';
blueBird.name;
```

# --instructions--

استخدم دالة البناء `Dog` من الدرس الأخير لإنشاء نسخة جديدة من `Dog`، مع تعيينها إلى متغير `hound`.

# --hints--

يجب إنشاء `hound` باستخدام دالة البناء `Dog`.

```js
assert(hound instanceof Dog);
```

يجب أن يستخدم شفرتك عامل `new` لإنشاء نسخة من `Dog`.

```js
assert(__helpers.removeJSComments(code).match(/new/g));
```

# --seed--

## --seed-contents--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
// Only change code below this line
```

# --solutions--

```js
function Dog() {
  this.name = "Rupert";
  this.color = "brown";
  this.numLegs = 4;
}
const hound = new Dog();
```
