---
id: 587d7db1367417b2b2512b85
title: ضبط النموذج الأولي للطفل ليكون نسخة من الأصل
challengeType: 1
forumTopicId: 301325
dashedName: set-the-childs-prototype-to-an-instance-of-the-parent
---

# --description--

في التحدٍّ السابق رأيت الخطوة الأولى لوراثة السلوك من النوع الأعلى (أو الأصل) `Animal`: إنشاء نسخة جديدة من `Animal`.

يغطي هذا التحدٍّ الخطوة التالية: ضبط `prototype` للنوع الفرعي (أو الطفل)—في هذه الحالة، `Bird`—ليكون نسخة من `Animal`.

```js
Bird.prototype = Object.create(Animal.prototype);
```

تذكّر أن `prototype` يشبه "الوصفة" لإنشاء كائن. بطريقة ما، تتضمن وصفة `Bird` الآن كل "المكونات" الأساسية من `Animal`.

```js
let duck = new Bird("Donald");
duck.eat();
```

يرث `duck` كل خصائص `Animal`، بما في ذلك طريقة `eat`.

# --instructions--

عدّل الشفرة بحيث ترث نسخ `Dog` من `Animal`.

# --hints--

يجب أن يكون `Dog.prototype` نسخة من `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }

// Only change code below this line


let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

function Dog() { }
Dog.prototype = Object.create(Animal.prototype);

let beagle = new Dog();
beagle.eat();
```
