---
id: 587d7dad367417b2b2512b75
title: إنشاء طريقة على كائن
challengeType: 1
forumTopicId: 301318
dashedName: create-a-method-on-an-object
---

# --description--

يمكن أن تحتوي الكائنات على نوع خاص من الخاصيات، يُسمى <dfn>طريقة</dfn>.

الطرق هي خاصيات تكون دوالًا. هذا يضيف سلوكًا مختلفًا للكائن. إليك مثال `duck` مع طريقة:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + duck.name + ".";}
};
duck.sayName();
```

يضيف المثال الطريقة `sayName`، وهي دالة تُرجع جملة تعطي اسم الـ `duck`. لاحظ أن الطريقة وصلت إلى الخاصية `name` في جملة الإرجاع باستخدام `duck.name`. التحدي التالي سيغطي طريقة أخرى للقيام بذلك.

# --instructions--

باستخدام كائن `dog`، أضف له طريقة تسمى `sayLegs`. يجب أن تُرجع الطريقة الجملة `This dog has 4 legs.`

# --hints--

يجب أن تكون `dog.sayLegs()` دالة.

```js
assert(typeof dog.sayLegs === 'function');
```

يجب أن تُرجع `dog.sayLegs()` السلسلة النصية المعطاة - لاحظ أن علامات الترقيم والمسافات مهمة.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,

};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
