---
id: 587d7db2367417b2b2512b8a
title: استخدام الإغلاق (closure) لحماية الخصائص داخل الكائن من التعديل الخارجي
  Use Closure to Protect Properties Within an Object from Being Modified
  Externally
challengeType: 1
forumTopicId: 18234
dashedName: >-
  use-closure-to-protect-properties-within-an-object-from-being-modified-externally
---

# --description--

في التحدي السابق، كان لدى `bird` خاصية عامة `name`. تُعتبر هذه الخاصية عامة لأنها يمكن الوصول إليها وتغييرها خارج تعريف `bird`.

```js
bird.name = "Duffy";
```

لذلك، يمكن لأي جزء من شفرتك تغيير اسم `bird` إلى أي قيمة بسهولة. فكر في أشياء مثل كلمات المرور والحسابات البنكية التي يمكن تغييرها بسهولة من أي جزء في قاعدة شفرتك. هذا قد يسبب الكثير من المشاكل.

أسهل طريقة لجعل هذه الخاصية العامة خاصة هي بإنشاء متغير داخل دالة البناء. هذا يغير نطاق ذلك المتغير ليكون داخل دالة البناء بدلاً من أن يكون متاحًا عالميًا. بهذه الطريقة، يمكن الوصول إلى المتغير وتغييره فقط بواسطة الطرق الموجودة أيضًا داخل دالة البناء.

```js
function Bird() {
  let hatchedEgg = 10;

  this.getHatchedEggCount = function() { 
    return hatchedEgg;
  };
}
let ducky = new Bird();
ducky.getHatchedEggCount();
```

هنا `getHatchedEggCount` هي طريقة مميزة، لأنها تملك وصولًا إلى المتغير الخاص `hatchedEgg`. هذا ممكن لأن `hatchedEgg` مُعلن في نفس السياق مع `getHatchedEggCount`. في جافاسكريبت، الدالة دائمًا تملك وصولًا إلى السياق الذي تم إنشاؤها فيه. هذا يسمى `closure`.

# --instructions--

غيّر طريقة إعلان `weight` في دالة `Bird` ليكون متغيرًا خاصًا. ثم أنشئ طريقة `getWeight` تُرجع قيمة `weight` 15.

# --hints--

يجب أن تكون خاصية `weight` متغيرًا خاصًا ويُعيّن لها القيمة `15`.

```js
assert(__helpers.removeJSComments(code).match(/(var|let|const)\s+weight\s*\=\s*15\;?/g));
```

يجب أن تنشئ شفرتك طريقة في `Bird` تسمى `getWeight` تُرجع قيمة المتغير الخاص `weight`.

```js
assert(new Bird().getWeight() === 15);
```

يجب أن تُرجع دالة `getWeight` المتغير الخاص `weight`.

```js
assert(__helpers.removeJSComments(code).match(/((return\s+)|(\(\s*\)\s*\=\>\s*))weight\;?/g));
```

# --seed--

## --seed-contents--

```js
function Bird() {
  this.weight = 15;


}
```

# --solutions--

```js
function Bird() {
  let weight = 15;

  this.getWeight = () => weight;
}
```
