---
id: 587d7dac367417b2b2512b74
title: استخدام ترميز النقطة للوصول إلى خصائص الكائن
challengeType: 1
forumTopicId: 301333
dashedName: use-dot-notation-to-access-the-properties-of-an-object
---

# --description--

التحدٍّ الأخير أنشأ كائنًا يحتوي على خصائص مختلفة. الآن سترى كيف تصل إلى قيم تلك الخصائص. إليك مثالًا:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
console.log(duck.name);
```

يُستخدم ترميز النقطة على اسم الكائن، `duck`، متبوعًا باسم الخاصية، `name`، للوصول إلى قيمة `Aflac`.

# --instructions--

اطبع خاصيتي كائن `dog` في وحدة التحكم.

# --hints--

يجب أن يستخدم شفرتك `console.log` لطباعة قيمة الخاصية `name` لكائن `dog`.

```js
assert(/console.log\(.*dog\.name.*\)/g.test(__helpers.removeJSComments(code)));
```

يجب أن يستخدم شفرتك `console.log` لطباعة قيمة الخاصية `numLegs` لكائن `dog`.

```js
assert(/console.log\(.*dog\.numLegs.*\)/g.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
// Only change code below this line
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4
};
console.log(dog.name);
console.log(dog.numLegs);
```
