---
id: 587d7daf367417b2b2512b7e
title: فهم خاصية دالة البناء
challengeType: 1
forumTopicId: 301327
dashedName: understand-the-constructor-property
---

# --description--

هناك خاصية `constructor` خاصة موجودة على مثيلات الكائنات `duck` و `beagle` التي تم إنشاؤها في التحديات السابقة:

```js
let duck = new Bird();
let beagle = new Dog();

console.log(duck.constructor === Bird); 
console.log(beagle.constructor === Dog);
```

كلا الاستدعاءين لـ `console.log` سيعرضان `true` في وحدة التحكم.

لاحظ أن خاصية `constructor` هي مرجع إلى دالة البناء التي أنشأت المثيل. ميزة خاصية `constructor` أنها تتيح التحقق من هذه الخاصية لمعرفة نوع الكائن. إليك مثالًا على كيفية استخدام ذلك:

```js
function joinBirdFraternity(candidate) {
  if (candidate.constructor === Bird) {
    return true;
  } else {
    return false;
  }
}
```

**ملاحظة:** بما أن خاصية `constructor` يمكن أن تُعاد كتابتها (وسيتم تغطية ذلك في التحديين التاليين) فمن الأفضل عمومًا استخدام طريقة `instanceof` للتحقق من نوع الكائن.

# --instructions--

اكتب دالة `joinDogFraternity` التي تأخذ معلمة `candidate` وتستخدم خاصية `constructor` لإرجاع `true` إذا كان المرشح من نوع `Dog`، وإلا ترجع `false`.

# --hints--

يجب تعريف `joinDogFraternity` كدالة.

```js
assert(typeof joinDogFraternity === 'function');
```

يجب أن ترجع `joinDogFraternity` القيمة `true` إذا كان `candidate` مثيلًا لـ `Dog`.

```js
assert(joinDogFraternity(new Dog('')) === true);
```

يجب أن تستخدم `joinDogFraternity` خاصية `constructor`.

```js
assert(/\.constructor/.test(__helpers.removeJSComments(code)) && !/instanceof/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
function joinDogFraternity(candidate) {

}
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
function joinDogFraternity(candidate) {
  return candidate.constructor === Dog;
}
```
