---
id: 587d7db2367417b2b2512b8c
title: استخدم تعبير دالة يتم استدعاؤه فورًا لإنشاء وحدة
challengeType: 1
forumTopicId: 301332
dashedName: use-an-iife-to-create-a-module
---

# --description--

غالبًا ما يُستخدم تعبير دالة يتم استدعاؤه فورًا (IIFE) لتجميع الوظائف المرتبطة في كائن واحد أو <dfn>وحدة</dfn>. على سبيل المثال، تحدٍّ سابق عرّف خليطين:

```js
function glideMixin(obj) {
  obj.glide = function() {
    console.log("Gliding on the water");
  };
}
function flyMixin(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  };
}
```

يمكننا تجميع هذه الخلائط في وحدة كما يلي:

```js
let motionModule = (function () {
  return {
    glideMixin: function(obj) {
      obj.glide = function() {
        console.log("Gliding on the water");
      };
    },
    flyMixin: function(obj) {
      obj.fly = function() {
        console.log("Flying, wooosh!");
      };
    }
  }
})();
```

لاحظ أن لديك تعبير دالة يتم استدعاؤه فورًا (IIFE) يُرجع كائنًا باسم `motionModule`. يحتوي هذا الكائن المرجع على كل سلوكيات الخليط كخصائص للكائن. ميزة نمط الوحدة هي أن كل سلوكيات الحركة يمكن تجميعها في كائن واحد يمكن استخدامه بعد ذلك في أجزاء أخرى من شفرتك. إليك مثالًا على استخدامه:

```js
motionModule.glideMixin(duck);
duck.glide();
```

# --instructions--

أنشئ وحدة باسم `funModule` لتغليف الخليطين `isCuteMixin` و `singMixin`. يجب أن تُرجع `funModule` كائنًا.

# --hints--

يجب تعريف `funModule` وأن تُرجع كائنًا.

```js
assert(typeof funModule === 'object');
```

يجب أن يصل `funModule.isCuteMixin` إلى دالة.

```js
assert(typeof funModule.isCuteMixin === 'function');
```

يجب أن يصل `funModule.singMixin` إلى دالة.

```js
assert(typeof funModule.singMixin === 'function');
```

# --seed--

## --seed-contents--

```js
let isCuteMixin = function(obj) {
  obj.isCute = function() {
    return true;
  };
};
let singMixin = function(obj) {
  obj.sing = function() {
    console.log("Singing to an awesome tune");
  };
};
```

# --solutions--

```js
const funModule = (function () {
  return {
    isCuteMixin: obj => {
      obj.isCute = () => true;
    },
    singMixin: obj => {
      obj.sing = () => console.log("Singing to an awesome tune");
    }
  };
})();
```
