---
id: 587d7db0367417b2b2512b82
title: فهم سلسلة النموذج الأولي
challengeType: 1
forumTopicId: 301329
dashedName: understand-the-prototype-chain
---

# --description--

كل الكائنات في JavaScript (مع بعض الاستثناءات) لها `prototype`. وأيضًا، `prototype` الخاص بالكائن نفسه هو كائن.

```js
function Bird(name) {
  this.name = name;
}

typeof Bird.prototype;
```

بما أن `prototype` هو كائن، يمكن أن يكون لـ `prototype` خاصته `prototype` أيضًا! في هذه الحالة، `prototype` الخاص بـ `Bird.prototype` هو `Object.prototype`:

```js
Object.prototype.isPrototypeOf(Bird.prototype);
```

كيف يكون هذا مفيدًا؟ قد تتذكر طريقة `hasOwnProperty` من تحدٍّ سابق:

```js
let duck = new Bird("Donald");
duck.hasOwnProperty("name");
```

تُعرَّف الطريقة `hasOwnProperty` في `Object.prototype`، التي يمكن الوصول إليها عبر `Bird.prototype`، التي يمكن بعدها الوصول إليها عبر `duck`. هذا مثال على سلسلة `prototype`. في سلسلة `prototype` هذه، تُعدّ `Bird` النوع الأعلى `supertype` لـ `duck`، بينما `duck` هي النوع الأدنى `subtype`. و`Object` نوع أعلى `supertype` لكلٍّ من `Bird` و`duck`. و`Object` نوع أعلى `supertype` لجميع الكائنات في JavaScript. لذلك، يمكن لأي كائن استخدام الطريقة `hasOwnProperty`.

# --instructions--

عدّل الشفرة لتُظهر سلسلة النموذج الأولي الصحيحة.

# --hints--

يجب أن تُظهر شفرتك أن `Object.prototype` هو النموذج الأولي لـ `Dog.prototype`

```js
assert(/Object\.prototype\.isPrototypeOf/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

Dog.prototype.isPrototypeOf(beagle);  // yields true

// Fix the code below so that it evaluates to true
???.isPrototypeOf(Dog.prototype);
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
Object.prototype.isPrototypeOf(Dog.prototype);
```
