---
id: 587d7db2367417b2b2512b8b
title: فهم التعبير عن الدالة التي تُنفذ فور إعلانها (IIFE)
challengeType: 1
forumTopicId: 301328
dashedName: understand-the-immediately-invoked-function-expression-iife
---

# --description--

نمط شائع في JavaScript هو تنفيذ دالة بمجرد إعلانها:

```js
(function () {
  console.log("Chirp, chirp!");
})();
```

هذا تعبير دالة مجهولة تُنفذ فورًا، وتُخرج `Chirp, chirp!` مباشرة.

لاحظ أن الدالة ليس لها اسم ولا تُخزن في متغير. الأقواس العادية () في نهاية التعبير تُجعل الدالة تُنفذ أو تُستدعى فورًا. يُعرف هذا النمط باسم <dfn>التعبير عن الدالة التي تُنفذ فور إعلانها</dfn> أو <dfn>IIFE</dfn>.

# --instructions--

أعد كتابة الدالة `makeNest` وأزل استدعاءها لتصبح بدلاً من ذلك تعبير دالة مجهولة تُنفذ فورًا (IIFE).

# --hints--

يجب أن تكون الدالة مجهولة.

```js
assert(/\((function|\(\))(=>|\(\)){?/.test(__helpers.removeJSComments(code).replace(/\s/g, '')));
```

يجب أن تحتوي دالتك على أقواس عادية في نهاية التعبير لاستدعائها فورًا.

```js
assert(/\(.*(\)\(|\}\(\))\)/.test(__helpers.removeJSComments(code).replace(/[\s;]/g, '')));
```

# --seed--

## --seed-contents--

```js
function makeNest() {
  console.log("A cozy nest is ready");
}

makeNest();
```

# --solutions--

```js
(function () {
  console.log("A cozy nest is ready");
})();
```

---

```js
(function () {
  console.log("A cozy nest is ready");
}());
```

---

```js
(() => {
  console.log("A cozy nest is ready");
})();
```

---

```js
(() =>
  console.log("A cozy nest is ready")
)();
```
