---
id: 587d7dae367417b2b2512b7b
title: فهم الخصائص الخاصة
challengeType: 1
forumTopicId: 301326
dashedName: understand-own-properties
---

# --description--

في المثال التالي، تُعرّف دالة البناء `Bird` خاصيتين: `name` و `numLegs`:

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let duck = new Bird("Donald");
let canary = new Bird("Tweety");
```

تُسمى `name` و `numLegs` <dfn>الخصائص الخاصة</dfn> لأنها تُعرّف مباشرة على كائن النسخة. هذا يعني أن كلًا من `duck` و `canary` له نسخة منفصلة خاصة به من هذه الخصائص. في الواقع، كل نسخة من `Bird` سيكون لها نسختها الخاصة من هذه الخصائص. الشفرة التالية تضيف كل الخصائص الخاصة بـ `duck` إلى المصفوفة `ownProps`:

```js
let ownProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  }
}

console.log(ownProps);
```

ستعرض وحدة التحكم القيمة `["name", "numLegs"]`.

# --instructions--

أضف الخصائص الخاصة بـ `canary` إلى المصفوفة `ownProps`.

# --hints--

يجب أن تتضمن `ownProps` القيم `numLegs` و `name`.

```js
assert(ownProps.indexOf('name') !== -1 && ownProps.indexOf('numLegs') !== -1);
```

يجب أن تحل هذا التحدٍّ دون استخدام الطريقة المدمجة `Object.keys()`.

```js
assert(!/Object(\.keys|\[(['"`])keys\2\])/.test(__helpers.removeJSComments(code)));
```

يجب أن تحل هذا التحدٍّ دون ترميز مصفوفة `ownProps` بشكل ثابت.

```js
assert(
  !/\[\s*(?:'|")(?:name|numLegs)|(?:push|concat)\(\s*(?:'|")(?:name|numLegs)/.test(
    __helpers.removeJSComments(code)
  )
);
```

# --seed--

## --seed-contents--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
let ownProps = [];
// Only change code below this line
```

# --solutions--

```js
function Bird(name) {
  this.name = name;
  this.numLegs = 2;
}

let canary = new Bird("Tweety");
function getOwnProps (obj) {
  const props = [];

  for (let prop in obj) {
    if (obj.hasOwnProperty(prop)) {
      props.push(prop);
    }
  }

  return props;
}

const ownProps = getOwnProps(canary);
```
