---
id: 587d7daf367417b2b2512b7f
title: تغيير النموذج الأولي إلى كائن جديد
challengeType: 1
forumTopicId: 301316
dashedName: change-the-prototype-to-a-new-object
---

# --description--

حتى الآن كنت تضيف الخصائص إلى الـ `prototype` بشكل فردي:

```js
Bird.prototype.numLegs = 2;
```

هذا يصبح مرهقًا بعد إضافة أكثر من بضع خصائص.

```js
Bird.prototype.eat = function() {
  console.log("nom nom nom");
}

Bird.prototype.describe = function() {
  console.log("My name is " + this.name);
}
```

طريقة أكثر كفاءة هي تعيين الـ `prototype` إلى كائن جديد يحتوي بالفعل على الخصائص. بهذه الطريقة، تُضاف الخصائص دفعة واحدة:

```js
Bird.prototype = {
  numLegs: 2, 
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --instructions--

أضف الخاصية `numLegs` والطريقتين `eat()` و `describe()` إلى الـ `prototype` الخاص بـ `Dog` عن طريق تعيين الـ `prototype` إلى كائن جديد.

# --hints--

يجب تعيين `Dog.prototype` إلى كائن جديد.

```js
assert(/Dog\.prototype\s*?=\s*?{/.test(__helpers.removeJSComments(code)));
```

يجب أن يحتوي `Dog.prototype` على الخاصية `numLegs`.

```js
assert(Dog.prototype.numLegs !== undefined);
```

يجب أن يحتوي `Dog.prototype` على الطريقة `eat()`.

```js
assert(typeof Dog.prototype.eat === 'function');
```

يجب أن يحتوي `Dog.prototype` على الطريقة `describe()`.

```js
assert(typeof Dog.prototype.describe === 'function');
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype = {
  // Only change code below this line

};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
numLegs: 4,
  eat () {
    console.log('nom nom nom');
  },
  describe () {
    console.log('My name is ' + this.name);
  }
};
```
