---
id: 587d7dad367417b2b2512b76
title: اجعل الشفرة أكثر قابلية لإعادة الاستخدام باستخدام الكلمة المفتاحية this
challengeType: 1
forumTopicId: 301321
dashedName: make-code-more-reusable-with-the-this-keyword
---

# --description--

التحدي الأخير قدّم طريقة إلى الكائن `duck`. استخدم ترميز النقطة `duck.name` للوصول إلى قيمة الخاصية `name` داخل جملة الإرجاع:

```js
sayName: function() {return "The name of this duck is " + duck.name + ".";}
```

بينما هذه طريقة صحيحة للوصول إلى خاصية الكائن، هناك مشكلة هنا. إذا تغير اسم المتغير، يجب تحديث أي شفرة تشير إلى الاسم الأصلي أيضًا. في تعريف كائن قصير، لا تشكل هذه مشكلة، لكن إذا كان للكائن العديد من المراجع لخصائصه، تزداد فرصة الخطأ.

طريقة لتجنب هذه المشاكل هي استخدام الكلمة المفتاحية `this`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2,
  sayName: function() {return "The name of this duck is " + this.name + ".";}
};
```

`this` موضوع عميق، والمثال أعلاه هو طريقة واحدة فقط لاستخدامه. في السياق الحالي، تشير `this` إلى الكائن المرتبطة به الطريقة: `duck`. إذا تغير اسم الكائن إلى `mallard`، لا حاجة للبحث عن كل المراجع لـ `duck` في الشفرة. هذا يجعل الشفرة قابلة لإعادة الاستخدام وأسهل في القراءة.

# --instructions--

عدّل الطريقة `dog.sayLegs` لإزالة أي مراجع إلى `dog`. استخدم مثال `duck` كدليل.

# --hints--

يجب أن تُرجع `dog.sayLegs()` السلسلة النصية المعطاة.

```js
assert(dog.sayLegs() === 'This dog has 4 legs.');
```

يجب أن تستخدم شفرتك الكلمة المفتاحية `this` للوصول إلى الخاصية `numLegs` في `dog`.

```js
assert(__helpers.removeJSComments(code).match(/this\.numLegs/g));
```

# --seed--

## --seed-contents--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs: function() {return "This dog has " + dog.numLegs + " legs.";}
};

dog.sayLegs();
```

# --solutions--

```js
let dog = {
  name: "Spot",
  numLegs: 4,
  sayLegs () {
    return 'This dog has ' + this.numLegs + ' legs.';
  }
};

dog.sayLegs();
```
