---
id: 587d7daf367417b2b2512b7d
title: التكرار على جميع الخصائص
challengeType: 1
forumTopicId: 301320
dashedName: iterate-over-all-properties
---

# --description--

لقد رأيت الآن نوعين من الخصائص: <dfn>الخصائص الذاتية</dfn> وخصائص `prototype`. تُعرَّف الخصائص الذاتية مباشرةً على كائن النسخة نفسه. أما خصائص prototype فتُعرَّف على `prototype`.

```js
function Bird(name) {
  this.name = name;  //own property
}

Bird.prototype.numLegs = 2; // prototype property

let duck = new Bird("Donald");
```

إليك كيف تضيف الخصائص الخاصة لـ `duck` إلى المصفوفة `ownProps` وخصائص الـ `prototype` إلى المصفوفة `prototypeProps`:

```js
let ownProps = [];
let prototypeProps = [];

for (let property in duck) {
  if(duck.hasOwnProperty(property)) {
    ownProps.push(property);
  } else {
    prototypeProps.push(property);
  }
}

console.log(ownProps);
console.log(prototypeProps);
```

سيعرض `console.log(ownProps)` في وحدة التحكم `["name"]`، وسيعرض `console.log(prototypeProps)` `["numLegs"]`.

# --instructions--

أضف جميع الخصائص الخاصة بـ `beagle` إلى المصفوفة `ownProps`. أضف جميع خصائص الـ `prototype` الخاصة بـ `Dog` إلى المصفوفة `prototypeProps`.

# --hints--

يجب أن تحتوي مصفوفة `ownProps` على `name` فقط.

```js
assert.deepEqual(ownProps, ['name']);
```

يجب أن تحتوي مصفوفة `prototypeProps` على `numLegs` فقط.

```js
assert.deepEqual(prototypeProps, ['numLegs']);
```

يجب أن تحل هذا التحدٍّ دون استخدام الطريقة المدمجة `Object.keys()`.

```js
assert(!/\Object.keys/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}

Dog.prototype.numLegs = 4;

let beagle = new Dog("Snoopy");

let ownProps = [];
let prototypeProps = [];
for (let prop in beagle) {
  if (beagle.hasOwnProperty(prop)) {
    ownProps.push(prop);
  } else {
    prototypeProps.push(prop);
  }
}
```
