---
id: 587d7dae367417b2b2512b7c
title: استخدام خصائص النموذج الأولي لتقليل الشفرة المكررة
challengeType: 1
forumTopicId: 301336
dashedName: use-prototype-properties-to-reduce-duplicate-code
---

# --description--

بما أن `numLegs` من المحتمل أن تكون لها نفس القيمة لكل مثيل من `Bird`، فبالتالي لديك متغير مكرر `numLegs` داخل كل مثيل من `Bird`.

قد لا تكون هذه مشكلة عندما يكون هناك مثيلان فقط، لكن تخيل لو كان هناك ملايين من المثيلات. سيكون ذلك عددًا كبيرًا من المتغيرات المكررة.

طريقة أفضل هي استخدام `prototype` الخاص بـ `Bird`. الخصائص في `prototype` تُشارك بين جميع نسخ `Bird`. إليك كيفية إضافة `numLegs` إلى `Bird prototype`:

```js
Bird.prototype.numLegs = 2;
```

الآن كل مثيلات `Bird` لديها الخاصية `numLegs`.

```js
console.log(duck.numLegs);
console.log(canary.numLegs);
```

بما أن كل المثيلات تمتلك تلقائيًا الخصائص الموجودة على `prototype`، فكر في `prototype` كـ"وصفة" لإنشاء الكائنات. لاحظ أن `prototype` الخاص بـ `duck` و `canary` هو جزء من دالة البناء `Bird` كـ `Bird.prototype`.

# --instructions--

أضف خاصية `numLegs` إلى `prototype` الخاص بـ `Dog`

# --hints--

يجب أن يكون لـ `beagle` خاصية `numLegs`.

```js
assert(beagle.numLegs !== undefined);
```

يجب أن تكون قيمة `beagle.numLegs` رقمًا.

```js
assert(typeof beagle.numLegs === 'number');
```

يجب أن تكون `numLegs` خاصية في `prototype` وليست خاصية مملوكة للمثيل.

```js
assert(beagle.hasOwnProperty('numLegs') === false);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}



// Only change code above this line
let beagle = new Dog("Snoopy");
```

# --solutions--

```js
function Dog (name) {
  this.name = name;
}
Dog.prototype.numLegs = 4;
let beagle = new Dog("Snoopy");
```
