---
id: 587d7db1367417b2b2512b86
title: إعادة تعيين خاصية دالة البناء الموروثة
challengeType: 1
forumTopicId: 301324
dashedName: reset-an-inherited-constructor-property
---

# --description--

عندما يرث كائن الـ `prototype` من كائن آخر، يرث أيضًا خاصية دالة البناء الخاصة بالنوع الأعلى.

إليك مثالًا:

```js
function Bird() { }
Bird.prototype = Object.create(Animal.prototype);
let duck = new Bird();
duck.constructor
```

لكن يجب أن يظهر كل من `duck` وكل مثيلات `Bird` أنها بُنيت بواسطة `Bird` وليس `Animal`. للقيام بذلك، يمكنك تعيين خاصية دالة البناء لـ `Bird` يدويًا إلى كائن `Bird`:

```js
Bird.prototype.constructor = Bird;
duck.constructor
```

# --instructions--

صحح الشفرة بحيث تُرجع `duck.constructor` و `beagle.constructor` دوال البناء الخاصة بهما.

# --hints--

يجب أن يكون `Bird.prototype` مثيلًا لـ `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Bird.prototype));
```

يجب أن تُرجع `duck.constructor` القيمة `Bird`.

```js
assert(duck.constructor === Bird);
```

يجب أن يكون `Dog.prototype` مثيلًا لـ `Animal`.

```js
assert(Animal.prototype.isPrototypeOf(Dog.prototype));
```

يجب أن تُرجع `beagle.constructor` القيمة `Dog`.

```js
assert(beagle.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Animal() { }
function Bird() { }
function Dog() { }

Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);

// Only change code below this line



let duck = new Bird();
let beagle = new Dog();
```

# --solutions--

```js
function Animal() { }
function Bird() { }
function Dog() { }
Bird.prototype = Object.create(Animal.prototype);
Dog.prototype = Object.create(Animal.prototype);
Dog.prototype.constructor = Dog;
Bird.prototype.constructor = Bird;
let duck = new Bird();
let beagle = new Dog();
```
