---
id: 587d7dae367417b2b2512b7a
title: التحقق من دالة البناء لكائن باستخدام instanceof
challengeType: 1
forumTopicId: 301337
dashedName: verify-an-objects-constructor-with-instanceof
---

# --description--

في كل مرة تنشئ فيها دالة البناء كائنًا جديدًا، يُقال إن ذلك الكائن هو <dfn>مثيل</dfn> لدالة البناء الخاصة به. توفر JavaScript طريقة سهلة للتحقق من ذلك باستخدام عامل `instanceof`. يسمح لك `instanceof` بمقارنة كائن مع دالة بناء، ويُرجع `true` أو `false` بناءً على ما إذا كان ذلك الكائن قد تم إنشاؤه بواسطة دالة البناء أم لا. إليك مثالًا:

```js
let Bird = function(name, color) {
  this.name = name;
  this.color = color;
  this.numLegs = 2;
}

let crow = new Bird("Alexis", "black");

crow instanceof Bird;
```

ستُرجع هذه الطريقة `instanceof` القيمة `true`.

إذا تم إنشاء كائن بدون استخدام دالة بناء، فإن `instanceof` سيتحقق من أنه ليس مثيلًا لتلك الدالة:

```js
let canary = {
  name: "Mildred",
  color: "Yellow",
  numLegs: 2
};

canary instanceof Bird;
```

ستُرجع هذه الطريقة `instanceof` القيمة `false`.

# --instructions--

أنشئ مثيلًا جديدًا من دالة البناء `House`، وسمّه `myHouse` مع تمرير عدد غرف النوم. ثم استخدم `instanceof` للتحقق من أنه مثيل لـ `House`.

# --hints--

يجب أن يحتوي `myHouse` على خاصية `numBedrooms` مُعيّنة إلى رقم.

```js
assert(typeof myHouse.numBedrooms === 'number');
```

يجب أن تتحقق من أن `myHouse` هو مثيل لـ `House` باستخدام عامل `instanceof`.

```js
assert(/myHouse\s*instanceof\s*House/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}

// Only change code below this line
```

# --solutions--

```js
function House(numBedrooms) {
  this.numBedrooms = numBedrooms;
}
const myHouse = new House(4);
console.log(myHouse instanceof House);
```
