---
id: 587d7db2367417b2b2512b89
title: استخدام خليط لإضافة سلوك مشترك بين كائنات غير مرتبطة
challengeType: 1
forumTopicId: 301331
dashedName: use-a-mixin-to-add-common-behavior-between-unrelated-objects
---

# --description--

كما رأيت، يُشارك السلوك عبر الوراثة. مع ذلك، هناك حالات لا تكون فيها الوراثة الحل الأفضل. الوراثة لا تعمل جيدًا مع الكائنات غير المرتبطة مثل `Bird` و `Airplane`. كلاهما يمكنه الطيران، لكن `Bird` ليس نوعًا من `Airplane` والعكس صحيح.

بالنسبة للكائنات غير المرتبطة، من الأفضل استخدام <dfn>الخليطات</dfn>. يسمح الخليط لكائنات أخرى باستخدام مجموعة من الدوال.

```js
let flyMixin = function(obj) {
  obj.fly = function() {
    console.log("Flying, wooosh!");
  }
};
```

يأخذ `flyMixin` أي كائن ويُعطيه طريقة `fly`.

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let plane = {
  model: "777",
  numPassengers: 524
};

flyMixin(bird);
flyMixin(plane);
```

هنا يُمرر `bird` و `plane` إلى `flyMixin`، الذي يعين بعد ذلك دالة `fly` لكل كائن. الآن يمكن لكل من `bird` و `plane` الطيران:

```js
bird.fly();
plane.fly();
```

ستعرض وحدة التحكم السلسلة النصية `Flying, wooosh!` مرتين، مرة لكل استدعاء `.fly()`.

لاحظ كيف يسمح الخليط بإعادة استخدام نفس طريقة `fly` من قبل الكائنات غير المرتبطة `bird` و `plane`.

# --instructions--

أنشئ خليطًا باسم `glideMixin` يُعرّف طريقة باسم `glide`. ثم استخدم `glideMixin` لمنح كل من `bird` و `boat` القدرة على الانزلاق.

# --hints--

يجب أن يُعلن شفرتك عن متغير `glideMixin` يكون دالة.

```js
assert(typeof glideMixin === 'function');
```

يجب أن تستخدم `glideMixin` على الكائن `bird` لمنحه طريقة `glide`.

```js
assert(typeof bird.glide === 'function');
```

يجب أن تستخدم `glideMixin` على الكائن `boat` لمنحه طريقة `glide`.

```js
assert(typeof boat.glide === 'function');
```

# --seed--

## --seed-contents--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};

// Only change code below this line
```

# --solutions--

```js
let bird = {
  name: "Donald",
  numLegs: 2
};

let boat = {
  name: "Warrior",
  type: "race-boat"
};
function glideMixin (obj) {
  obj.glide = () => 'Gliding!';
}

glideMixin(bird);
glideMixin(boat);
```
