---
id: 587d7dac367417b2b2512b73
title: إنشاء كائن JavaScript أساسي
challengeType: 1
forumTopicId: 301317
dashedName: create-a-basic-javascript-object
---

# --description--

فكر في الأشياء التي يراها الناس يوميًا، مثل السيارات والمتاجر والطيور. كلها <dfn>كائنات</dfn>: أشياء ملموسة يمكن للناس ملاحظتها والتفاعل معها.

ما بعض صفات هذه الكائنات؟ السيارة لها عجلات. المتاجر تبيع سلعًا. الطيور لها أجنحة.

هذه الصفات، أو <dfn>خصائص</dfn>، تحدد ما يتكون منه الكائن. لاحظ أن الكائنات المتشابهة تشترك في نفس الخصائص، لكن قد تختلف في القيم الخاصة بتلك الخصائص. على سبيل المثال، كل السيارات لها عجلات، لكن ليس كل السيارات لها نفس عدد العجلات.

تُستخدم الكائنات في JavaScript لنمذجة الكائنات الواقعية، مع إعطائها خصائص وسلوكًا مثل نظيراتها في العالم الحقيقي. إليك مثالًا يستخدم هذه المفاهيم لإنشاء كائن `duck`:

```js
let duck = {
  name: "Aflac",
  numLegs: 2
};
```

هذا الكائن `duck` يحتوي على زوجين من الخاصية/القيمة: خاصية `name` بقيمة `Aflac` و`numLegs` بقيمة 2.

# --instructions--

أنشئ كائن `dog` بخصائص `name` و`numLegs`، واضبطهما على سلسلة نصية ورقم، على التوالي.

# --hints--

يجب أن يكون `dog` كائنًا.

```js
assert(typeof dog === 'object');
```

يجب أن يحتوي `dog` على خاصية `name` مضبوطة على سلسلة نصية.

```js
assert(typeof dog.name === 'string');
```

يجب أن يحتوي `dog` على خاصية `numLegs` مضبوطة على رقم.

```js
assert(typeof dog.numLegs === 'number');
```

# --seed--

## --seed-contents--

```js
let dog = {

};
```

# --solutions--

```js
let dog = {
  name: '',
  numLegs: 4
};
```
