---
id: 587d7db0367417b2b2512b81
title: فهم مصدر النموذج الأولي للكائن
challengeType: 1
forumTopicId: 301330
dashedName: understand-where-an-objects-prototype-comes-from
---

# --description--

تمامًا كما يرث الأشخاص الجينات من والديهم، يرث الكائن النموذج الأولي (`prototype`) مباشرة من دالة البناء التي أنشأته. على سبيل المثال، هنا دالة البناء `Bird` تنشئ الكائن `duck`:

```js
function Bird(name) {
  this.name = name;
}

let duck = new Bird("Donald");
```

يرث `duck` النموذج الأولي (`prototype`) من دالة البناء `Bird`. يمكنك إظهار هذه العلاقة باستخدام الطريقة `isPrototypeOf`:

```js
Bird.prototype.isPrototypeOf(duck);
```

سيُرجع هذا القيمة `true`.

# --instructions--

استخدم `isPrototypeOf` للتحقق من النموذج الأولي (`prototype`) الخاص بـ `beagle`.

# --hints--

يجب أن تُظهر أن `Dog.prototype` هو النموذج الأولي (`prototype`) الخاص بـ `beagle`

```js
assert(/Dog\.prototype\.isPrototypeOf\(beagle\)/.test(__helpers.removeJSComments(code)));
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

let beagle = new Dog("Snoopy");

// Only change code below this line
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
let beagle = new Dog("Snoopy");
Dog.prototype.isPrototypeOf(beagle);
```
