---
id: 587d7daf367417b2b2512b80
title: تذكّر تعيين خاصية دالة البناء عند تغيير النموذج الأولي
challengeType: 1
forumTopicId: 301323
dashedName: remember-to-set-the-constructor-property-when-changing-the-prototype
---

# --description--

هناك أثر جانبي مهم عند تعيين النموذج الأولي يدويًا إلى كائن جديد. إذ يمحو خاصية `constructor`! يمكن استخدام هذه الخاصية للتحقق من دالة البناء التي أنشأت النسخة، ولكن بما أن الخاصية قد تم استبدالها، فإنها تعطي نتائج خاطئة الآن:

```js
duck.constructor === Bird;
duck.constructor === Object;
duck instanceof Bird;
```

بالترتيب، ستُقيّم هذه التعبيرات إلى `false`، و`true`، و`true`.

لإصلاح ذلك، كلما تم تعيين نموذج أولي يدويًا إلى كائن جديد، تذكّر تعريف خاصية `constructor`:

```js
Bird.prototype = {
  constructor: Bird,
  numLegs: 2,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name); 
  }
};
```

# --instructions--

عرّف خاصية `constructor` على `prototype` الخاص بـ `Dog`.

# --hints--

يجب أن يعيّن `Dog.prototype` خاصية `constructor`.

```js
assert(Dog.prototype.constructor === Dog);
```

# --seed--

## --seed-contents--

```js
function Dog(name) {
  this.name = name;
}

// Only change code below this line
Dog.prototype = {

  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```

# --solutions--

```js
function Dog(name) {
  this.name = name;
}
Dog.prototype = {
  constructor: Dog,
  numLegs: 4,
  eat: function() {
    console.log("nom nom nom");
  },
  describe: function() {
    console.log("My name is " + this.name);
  }
};
```
