---
id: 587d7dad367417b2b2512b77
title: تعريف دالة البناء
challengeType: 1
forumTopicId: 16804
dashedName: define-a-constructor-function
---

# --description--

<dfn>دوال البناء</dfn> هي دوال تُنشئ كائنات جديدة. تُعرّف الخصائص والسلوكيات التي ستنتمي إلى الكائن الجديد. فكر فيها كأنها مخطط لإنشاء كائنات جديدة.

إليك مثالًا على دالة بناء:

```js
function Bird() {
  this.name = "Albert";
  this.color = "blue";
  this.numLegs = 2;
}
```

تعرف دالة البناء هذه كائن `Bird` بخواص `name` و`color` و`numLegs` التي تم تعيينها إلى Albert وblue و2 على التوالي. تتبع دوال البناء بعض الاصطلاحات:

<ul><li>تُعرف دوال البناء باسم يبدأ بحرف كبير لتمييزها عن الدوال الأخرى التي ليست <code>constructors</code>.</li><li>تستخدم دوال البناء الكلمة المفتاحية <code>this</code> لتعيين خصائص الكائن الذي ستنشئه. داخل دالة البناء، تشير <code>this</code> إلى الكائن الجديد الذي ستنشئه.</li><li>تعرف دوال البناء الخصائص والسلوكيات بدلاً من إرجاع قيمة كما تفعل الدوال الأخرى.</li></ul>

# --instructions--

أنشئ دالة بناء، `Dog`، بخواص `name` و`color` و`numLegs` التي تُعيّن إلى سلسلة نصية، وسلسلة نصية، ورقم على التوالي.

# --hints--

يجب أن تحتوي `Dog` على خاصية `name` مُعيّنة إلى سلسلة نصية.

```js
assert(typeof new Dog().name === 'string');
```

يجب أن تحتوي `Dog` على خاصية `color` مُعيّنة إلى سلسلة نصية.

```js
assert(typeof new Dog().color === 'string');
```

يجب أن تحتوي `Dog` على خاصية `numLegs` مُعيّنة إلى رقم.

```js
assert(typeof new Dog().numLegs === 'number');
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function Dog (name, color, numLegs) {
  this.name = 'name';
  this.color = 'color';
  this.numLegs = 4;
}
```
