---
id: 587d7db0367417b2b2512b84
title: وراثة السلوكيات من نوع أعلى
challengeType: 1
forumTopicId: 301319
dashedName: inherit-behaviors-from-a-supertype
---

# --description--

في التحدٍّ السابق، أنشأت `supertype` يسمى `Animal` عرّف سلوكيات مشتركة بين جميع الحيوانات:

```js
function Animal() { }
Animal.prototype.eat = function() {
  console.log("nom nom nom");
};
```

يغطي هذا التحدٍّ والتحدٍّ التالي كيفية إعادة استخدام طرق `Animal` داخل `Bird` و `Dog` دون تعريفها مجددًا. يستخدم تقنية تسمى الوراثة. يغطي هذا التحدٍّ الخطوة الأولى: إنشاء نسخة من الـ `supertype` (أو الأصل). أنت تعرف بالفعل طريقة واحدة لإنشاء نسخة من `Animal` باستخدام عامل `new`:

```js
let animal = new Animal();
```

هناك بعض العيوب عند استخدام هذه الصيغة للوراثة، وهي معقدة جدًا بالنسبة لنطاق هذا التحدٍّ. بدلًا من ذلك، إليك طريقة بديلة بدون تلك العيوب:

```js
let animal = Object.create(Animal.prototype);
```

`Object.create(obj)` ينشئ كائنًا جديدًا، ويعيّن `obj` كـ `prototype` للكائن الجديد. تذكر أن الـ `prototype` يشبه "الوصفة" لإنشاء كائن. من خلال تعيين `prototype` الخاص بـ `animal` ليكون `prototype` الخاص بـ `Animal`، فإنك تعطي نسخة `animal` نفس "الوصفة" مثل أي نسخة أخرى من `Animal`.

```js
animal.eat();
animal instanceof Animal;
```

طريقة `instanceof` هنا سترجع `true`.

# --instructions--

استخدم `Object.create` لإنشاء نسختين من `Animal` باسم `duck` و `beagle`.

# --hints--

يجب تعريف المتغير `duck`.

```js
assert(typeof duck !== 'undefined');
```

يجب تعريف المتغير `beagle`.

```js
assert(typeof beagle !== 'undefined');
```

يجب تهيئة المتغير `duck` باستخدام `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}duck\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

يجب تهيئة المتغير `beagle` باستخدام `Object.create`.

```js
assert(
  /(let|const|var)\s{1,}beagle\s*=\s*Object\.create\s*\(\s*Animal\.prototype\s*\)\s*/.test(
    __helpers.removeJSComments(code)
  )
);
```

يجب أن يكون لـ `duck` `prototype` من نوع `Animal`.

```js
assert(duck instanceof Animal);
```

يجب أن يكون لـ `beagle` `prototype` من نوع `Animal`.

```js
assert(beagle instanceof Animal);
```

# --seed--

## --seed-contents--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};

// Only change code below this line

let duck; // Change this line
let beagle; // Change this line
```

# --solutions--

```js
function Animal() { }

Animal.prototype = {
  constructor: Animal,
  eat: function() {
    console.log("nom nom nom");
  }
};
let duck = Object.create(Animal.prototype);
let beagle = Object.create(Animal.prototype);

duck.eat();
beagle.eat();
```
