---
id: 587d7db1367417b2b2512b88
title: تجاوز الطرق الموروثة
challengeType: 1
forumTopicId: 301322
dashedName: override-inherited-methods
---

# --description--

في الدروس السابقة، تعلمت أن الكائن يمكنه وراثة سلوكه (الطرق) من كائن آخر عبر الإشارة إلى كائن `prototype` الخاص به:

```js
ChildObject.prototype = Object.create(ParentObject.prototype);
```

ثم حصل `ChildObject` على طرقه الخاصة بتسلسلها على `prototype` الخاص به:

```js
ChildObject.prototype.methodName = function() {...};
```

يمكنك تجاوز طريقة موروثة. يتم ذلك بنفس الطريقة - بإضافة طريقة إلى `ChildObject.prototype` باستخدام نفس اسم الطريقة التي تريد تجاوزها. إليك مثالًا على `Bird` يتجاوز طريقة `eat()` الموروثة من `Animal`:

```js
function Animal() { }
Animal.prototype.eat = function() {
  return "nom nom nom";
};
function Bird() { }

Bird.prototype = Object.create(Animal.prototype);

Bird.prototype.eat = function() {
  return "peck peck peck";
};
```

إذا كان لديك نسخة `let duck = new Bird();` واستدعيت `duck.eat()`، فهذه هي الطريقة التي يبحث بها JavaScript عن الطريقة في سلسلة `prototype` الخاصة بـ `duck`:

1.  `duck` => هل تم تعريف `eat()` هنا؟ لا.
2.  `Bird` => هل تم تعريف `eat()` هنا؟ => نعم. نفذها وتوقف عن البحث.
3.  `Animal` => تم تعريف `eat()` أيضًا، لكن JavaScript توقف عن البحث قبل الوصول إلى هذا المستوى.
4.  Object => توقف JavaScript عن البحث قبل الوصول إلى هذا المستوى.

# --instructions--

تجاوز طريقة `fly()` لـ `Penguin` بحيث تُرجع السلسلة النصية `Alas, this is a flightless bird.`

# --hints--

يجب أن تُرجع `penguin.fly()` السلسلة النصية `Alas, this is a flightless bird.`

```js
assert(penguin.fly() === 'Alas, this is a flightless bird.');
```

يجب أن تُرجع طريقة `bird.fly()` السلسلة النصية `I am flying!`

```js
assert(new Bird().fly() === 'I am flying!');
```

# --seed--

## --seed-contents--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;

// Only change code below this line



// Only change code above this line

let penguin = new Penguin();
console.log(penguin.fly());
```

# --solutions--

```js
function Bird() { }

Bird.prototype.fly = function() { return "I am flying!"; };

function Penguin() { }
Penguin.prototype = Object.create(Bird.prototype);
Penguin.prototype.constructor = Penguin;
Penguin.prototype.fly = () => 'Alas, this is a flightless bird.';
let penguin = new Penguin();
console.log(penguin.fly());
```
