---
id: 68e5293bd00d2fe134f58982
title: Schritt 9
challengeType: 20
dashedName: step-9
---

# --description--

Wie in der vorherigen Lektion besprochen, sind abstrakte Methoden Methoden, die von jeder Klasse implementiert werden mĂĽssen, die von der abstrakten Basisklasse erbt. Sie dienen als Vertrag und stellen sicher, dass alle Rabattstrategien bestimmte erforderliche Methoden enthalten.

Definieren Sie innerhalb Ihrer `DiscountStrategy`-Klasse eine Methode mit dem Namen `is_applicable`. Verwenden Sie den `@abstractmethod`-Dekorator ĂĽber der Methodendefinition. Die Methode sollte die folgenden Parameter haben: `self`, `product` und `user_tier`. Verwenden Sie die `pass`-Anweisung im Methodenrumpf.

FĂĽgen Sie Typangaben hinzu, um die Methodensignatur klar zu machen: `product` sollte vom Typ `Product` sein, `user_tier` sollte vom Typ `str` sein.


# --hints--

Ihre `DiscountStrategy`-Klasse sollte eine Methode mit dem Namen `is_applicable` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

Ihre `is_applicable`-Methode sollte den `@abstractmethod`-Dekorator haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

Ihre `is_applicable`-Methode sollte die Parameter `self`, `product: Product` und `user_tier: str` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
--fcc-editable-region--
    pass
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
