---
id: 68e529226b0defdff26c4a8a
title: Schritt 6
challengeType: 20
dashedName: step-6
---

# --description--

Erstellen Sie nun eine Instanz von `Product` mit dem Namen `Wireless Mouse` und einem Preis von `50.0`. Weisen Sie diese einer Variablen namens `product` zu. Geben Sie dann `product` in der Konsole aus, um zu sehen, wie es aussieht.

# --hints--

Sie sollten eine Variable namens `product` haben.

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

Sie sollten eine Instanz von `Product` der Variablen `product` zuweisen.

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

Ihre `Product`-Instanz sollte den Namen `Wireless Mouse` und den Preis `50.0` haben.

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

Sie sollten `product` in der Konsole ausgeben.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

Wenn Sie eine `Product` mit dem Namen `'Wireless Mouse'` und dem Preis `50.0` ausgeben, sollte `Wireless Mouse - $50.0` angezeigt werden.

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--

--fcc-editable-region--
```
