---
id: 68e5293bd00d2fe134f58988
title: Schritt 17
challengeType: 20
dashedName: step-17
---

# --description--

Implementieren Sie die `is_applicable`-Methode für Ihre `FixedAmountDiscount`-Klasse. Ein Festbetragsrabatt sollte nur angewendet werden, wenn es wirtschaftlich sinnvoll ist.

Geben Sie `True` zurück, wenn `product.price * 0.9` größer als `self.amount` ist, andernfalls geben Sie `False` zurück. So wird sichergestellt, dass der Rabatt weniger als 90 % des Originalpreises beträgt und somit realistisch bleibt.

# --hints--

Ihre `FixedAmountDiscount`-Klasse sollte eine `is_applicable`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("is_applicable")`)
})
```

Ihre `is_applicable`-Methode sollte die Parameter `self`, `product: Product` und `user_tier: str` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_args("self, product:Product, user_tier:str")`)
})
```

Ihre `is_applicable`-Methode sollte einen Rückgabetyp-Hinweis von `-> bool` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

Ihre `is_applicable`-Methode sollte `True` zurückgeben, wenn `product.price * 0.9` größer als `self.amount` ist.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(50)
    assert discount.is_applicable(test_product, "regular") == True
  `)
})
```

Ihre `is_applicable`-Methode sollte `False` zurückgeben, wenn `product.price * 0.9` kleiner oder gleich `self.amount` ist.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestFixedAmountDiscount(FixedAmountDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestFixedAmountDiscount(95)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
