---
id: 68e5293bd00d2fe134f58987
title: Schritt 16
challengeType: 20
dashedName: step-16
---

# --description--

Jetzt erstellen Sie eine weitere Rabattstrategie. Erstellen Sie eine Klasse namens `FixedAmountDiscount`, die von `DiscountStrategy` erbt.

Geben Sie ihr eine `__init__`-Methode, die `self` und `amount` als Parameter erhält. Geben Sie dem Parameter `amount` einen Typ-Hinweis `int` und der Methode einen Rückgabetyp-Hinweis `None`. Speichern Sie innerhalb der `__init__`-Methode `amount` als Instanzattribut.

# --hints--

Sie sollten eine Klasse namens `FixedAmountDiscount` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

Ihre `FixedAmountDiscount`-Klasse sollte von `DiscountStrategy` erben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

Ihre `FixedAmountDiscount`-Klasse sollte eine `__init__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

Ihre `__init__`-Methode sollte die Parameter `self` und `amount: int` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

Ihre `__init__`-Methode sollte einen Rückgabetyp von `None` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

Sie sollten `amount` innerhalb Ihrer `self.amount`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--

--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
