---
id: 68e529226b0defdff26c4a8b
title: Schritt 7
challengeType: 20
dashedName: step-7
---

# --description--

Jetzt beginnen Sie mit dem Aufbau der Schnittstelle für das Rabattsystem. Python bietet eine Möglichkeit, abstrakte Klassen zu definieren, die als Vorlage für andere Klassen dienen. Dies geschieht mit der `ABC`-Klasse (Abstract Base Class) des `abc`-Moduls.

Fügen Sie am Anfang Ihres Codes die folgende Import-Anweisung hinzu:

```py
from abc import ABC, abstractmethod
```

Damit importieren Sie die `ABC`-Klasse und den `abstractmethod`-Dekorator, die Sie verwenden, um Methoden zu definieren, die von jeder Klasse, die von Ihrer abstrakten Klasse erbt, implementiert werden müssen.

# --hints--

Sie sollten die erwähnte Import-Anweisung am Anfang Ihres Codes haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
