---
id: 68e5293bd00d2fe134f58985
title: Schritt 13
challengeType: 20
dashedName: step-13
---

# --description--

Da `PercentageDiscount` von einer abstrakten Klasse erbt, müssen Sie alle abstrakten Methoden implementieren. Beginnen Sie mit der Implementierung der `is_applicable`-Methode. Die `is_applicable`-Methode sollte `self`, `product: Product` und `user_tier: str` als Parameter erhalten. Sie sollte einen Rückgabetyp-Hinweis von `bool` haben.

Diese Methode soll bestimmen, ob ein prozentualer Rabatt gültig ist. Geben Sie `True` zurück, wenn `self.percent` kleiner oder gleich `70` ist, andernfalls geben Sie `False` zurück. So werden unrealistische Rabatte wie 90 % Rabatt verhindert.

# --hints--

Ihre `PercentageDiscount`-Klasse sollte eine `is_applicable`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

Ihre `is_applicable`-Methode sollte die Parameter `self`, `product: Product` und `user_tier: str` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

Ihre `is_applicable`-Methode sollte `True` zurückgeben, wenn `self.percent` kleiner oder gleich `70` ist.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

Ihre `is_applicable`-Methode sollte `False` zurückgeben, wenn `self.percent` größer als `70` ist.

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

Ihre `is_applicable`-Methode sollte einen Rückgabetyp-Hinweis von `bool` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

--fcc-editable-region--
    
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
