---
id: 68e52994bc7ea2e3dec5309d
title: Schritt 20
challengeType: 20
dashedName: step-20
---

# --description--

Für die dritte Rabattstrategie erstellen Sie eine Klasse namens `PremiumUserDiscount`, die von `DiscountStrategy` erbt.

Im Gegensatz zu den vorherigen Strategien benötigt diese keine `__init__`-Methode, da sie keine Konfiguration speichert. Sie wird allen Premium-Nutzern einen festen Rabatt von 20 % gewähren.

Erstellen Sie vorerst nur die leere Klasse, die von `DiscountStrategy` mit `pass` erbt.

# --hints--

Sie sollten eine Klasse namens `PremiumUserDiscount` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PremiumUserDiscount")`)
})
```

Ihre `PremiumUserDiscount`-Klasse sollte von `DiscountStrategy` erben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PremiumUserDiscount").inherits_from("DiscountStrategy")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
