---
id: 68e5293bd00d2fe134f58986
title: Schritt 14
challengeType: 20
dashedName: step-14
---

# --description--

Implementieren Sie nun die `apply_discount`-Methode in Ihrer `PercentageDiscount`.

Diese Methode soll den rabattierten Preis berechnen und zurückgeben.

Die Formel lautet: `product.price * (1 - self.percent / 100)`

# --hints--

Ihre `PercentageDiscount`-Klasse sollte eine `apply_discount`-Methode enthalten.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

Ihre `apply_discount`-Methode sollte die Parameter `self` und `product: Product` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

Ihre `apply_discount`-Methode sollte einen Rückgabetyp-Hinweis von `float` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

Ihre `apply_discount`-Methode sollte den rabattierten Preis mit der Formel `product.price * (1 - self.percent / 100)` berechnen.

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

--fcc-editable-region--
    
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
