---
id: 68e5293bd00d2fe134f58981
title: Schritt 8
challengeType: 20
dashedName: step-8
---

# --description--

Erstellen Sie eine Klasse namens `DiscountStrategy`, die von `ABC` erbt. Dies wird eine abstrakte Basisklasse sein, die das Interface für alle Rabattstrategien definiert.

Verwenden Sie vorerst einfach `pass` im Klassenkörper, um eine leere Klasse zu erstellen.

# --hints--

Sie sollten eine Klasse namens `DiscountStrategy` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

Ihre `DiscountStrategy`-Klasse sollte von `ABC` erben.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
