---
id: 68e5293bd00d2fe134f58984
title: Schritt 12
challengeType: 20
dashedName: step-12
---

# --description--

Jetzt ist es Zeit, eine konkrete Implementierung Ihrer abstrakten Klasse zu erstellen. Erstellen Sie eine Klasse namens `PercentageDiscount`, die von `DiscountStrategy` erbt.

Geben Sie ihr eine `__init__`-Methode, die `self` und `percent` als Parameter entgegennimmt. Geben Sie für `int` einen Typ-Hinweis `percent` an und setzen Sie den Rückgabetyp der Methode auf `None`. Speichern Sie innerhalb der `__init__`-Methode `percent` als Instanzattribut.

# --hints--

Sie sollten eine Klasse namens `PercentageDiscount` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

Ihre `PercentageDiscount`-Klasse sollte von `DiscountStrategy` erben.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

Ihre `PercentageDiscount`-Klasse sollte eine `__init__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

Ihre `__init__`-Methode sollte die Parameter `self` und `percent: int` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

Ihre `__init__`-Methode sollte einen Rückgabetyp-Hinweis `None` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

Sie sollten `percent` innerhalb Ihrer `self.percent`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
