---
id: 68e52994bc7ea2e3dec530a6
title: Schritt 29
challengeType: 20
dashedName: step-29
---

# --description--

Es ist Zeit, das komplette Rabatt-System zu testen! Fügen Sie eine `if`-Anweisung hinzu, die überprüft, ob `__name__ == '__main__'` ist, und verschieben Sie `product = Product('Wireless Mouse', 50.0)` hinein.

Erstellen Sie danach innerhalb des `if`-Blocks eine Variable `user_tier` und setzen Sie sie auf den String `Premium`. Erstellen Sie dann eine Variable `strategies`, die eine Liste von Strategien enthält, darunter `PercentageDiscount(10)`, `FixedAmountDiscount(5)` und `PremiumUserDiscount()`.

# --hints--

Sie sollten eine `if`-Anweisung haben, die überprüft, ob `__name__ == '__main__'` ist.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent('__name__ == "__main__"')`)
})
```

Sie sollten `product = Product('Wireless Mouse', 50.0)` innerhalb des `if`-Blocks verschieben.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt("product = Product('Wireless Mouse', 50.0)")`)
})
```

Sie sollten `user_tier` innerhalb des `Premium`-Blocks auf den String `if` setzen.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('user_tier = "Premium"')`)
})
```

Sie sollten eine Liste namens `strategies` erstellen, die `PercentageDiscount(10)`, `FixedAmountDiscount(5)` und `PremiumUserDiscount()` enthält.

```js
({
  test: () => runPython(`assert _Node(_code).find_ifs()[0].has_stmt('strategies = [PercentageDiscount(10), FixedAmountDiscount(5), PremiumUserDiscount()]')`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies

    def calculate_best_price(self, product: Product, user_tier: str) -> float:
        prices = [product.price]

        for strategy in self.strategies:
            if strategy.is_applicable(product, user_tier):
                discounted = strategy.apply_discount(product)
                prices.append(discounted)
        return min(prices)

--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
--fcc-editable-region--
```
