---
id: 68e529226b0defdff26c4a89
title: Schritt 3
challengeType: 20
dashedName: step-3
---

# --description--

Im Moment sehen Sie, wenn Sie eine Instanz von `Product` erstellen und ausgeben, eine Standardobjektdarstellung, die nicht sehr aussagekräftig ist. Zum Beispiel könnte sie ungefähr so aussehen: `<__main__.Product object at 0x7f9b8c2d1d30>`.

Um die Standardobjektdarstellung in einen lesbaren String umzuwandeln, erstellen Sie eine `__str__`-Methode in Ihrer `Product`-Klasse, die einen String im Format `name - $price` zurückgibt (wobei `name` und `price` durch die Werte der entsprechenden Attribute ersetzt werden sollten).

# --hints--

Ihre `Product`-Klasse sollte eine `__str__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

Ihre `__str__`-Methode sollte einen String im Format `name - $price` zurückgeben.

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
        
--fcc-editable-region--
    
--fcc-editable-region--
```
