---
id: 68e52994bc7ea2e3dec530a1
title: Schritt 23
challengeType: 20
dashedName: step-23
---

# --description--

Jetzt benötigen Sie eine Möglichkeit, mehrere Rabattstrategien zu verwalten und den besten Preis zu ermitteln.

Erstellen Sie eine Klasse namens `DiscountEngine`. Diese Klasse verwaltet alle Rabattstrategien und berechnet den besten Preis für ein Produkt.

Geben Sie ihr eine `__init__`-Methode, die `self` und `strategies` als Parameter erhält, wobei `strategies` einen Typ-Hinweis von `list[DiscountStrategy]` hat. Die Verwendung von `list[DiscountStrategy]` gibt an, dass der Parameter eine Liste von `DiscountStrategy`-Objekten ist. Der Rückgabewert der Methode sollte `None` sein. Speichern Sie innerhalb der `__init__`-Methode `strategies` als Instanzattribut.

# --hints--

Sie sollten eine Klasse namens `DiscountEngine` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

Ihre `DiscountEngine`-Klasse sollte eine `__init__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

Ihre `__init__`-Methode sollte die Parameter `self` und `strategies: list[DiscountStrategy]` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

Ihre `__init__`-Methode sollte einen Rückgabetyp von `None` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

Sie sollten `strategies` in Ihrer `self.strategies`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
