---
id: 68e5293bd00d2fe134f58983
title: Schritt 11
challengeType: 20
dashedName: step-11
---

# --description--

Fügen Sie nun eine zweite abstrakte Methode namens `apply_discount` zu Ihrer `DiscountStrategy`-Klasse hinzu. Verwenden Sie den `@abstractmethod`-Decorator und geben Sie ihm die Parameter `self` und `product`.

Fügen Sie Typangaben hinzu: `product` sollte vom Typ `Product` sein, und die Methode sollte einen `float` zurückgeben. Verwenden Sie `pass` im Methodenrumpf.

# --hints--

Ihre `DiscountStrategy`-Klasse sollte eine Methode namens `apply_discount` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

Ihre `apply_discount`-Methode sollte den `@abstractmethod`-Decorator haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

Ihre `apply_discount`-Methode sollte die Parameter `self` und `product: Product` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

Ihre `apply_discount`-Methode sollte einen Rückgabetyp-Hinweis von `-> float` haben.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
        
--fcc-editable-region--
    
--fcc-editable-region--

product = Product('Wireless Mouse', 50.0)
print(product)
```
