---
id: 68e524c94e2e4ab7b4bfc7e2
title: Schritt 1
challengeType: 20
dashedName: step-1
---

# --description--

In diesem Workshop bauen Sie einen Rabattrechner, der verschiedene Rabattstrategien auf Produkte anwenden kann. Das System ermittelt den besten Preis für einen Kunden basierend auf mehreren Rabattregeln.

Beginnen Sie damit, eine Klasse mit dem Namen `Product` zu definieren. Geben Sie Ihrer Klasse eine `__init__`-Methode mit den folgenden Parametern: `self`, `name` und `price`. Weisen Sie innerhalb der `__init__`-Methode die Parameter `name` und `price` Instanzattributen mit demselben Namen zu.

# --hints--

Sie sollten eine Klasse mit dem Namen `Product` erstellen.

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

Ihre `Product`-Klasse sollte eine `__init__`-Methode haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

Ihre `__init__`-Methode sollte die folgenden Parameter haben: `self`, `name` und `price`.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

Sie sollten `name` in Ihrer `self.name`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

Sie sollten `price` in Ihrer `self.price`-Methode `__init__` zuweisen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
