---
id: 68e79d91daf8fb4808f93c47
title: Schritt 15
challengeType: 20
dashedName: step-15
---

# --description--

Nachdem Sie die `apply_discount`-Methode implementiert haben, testen Sie Ihre Rabattstrategie. Erstellen Sie eine Instanz von `PercentageDiscount` mit `10` Prozent und weisen Sie sie einer Variablen namens `discount` zu. Wenden Sie dann den Rabatt mit `discount.apply_discount(product)` an und geben Sie das Ergebnis aus, um den Rabatt in Aktion zu sehen.

# --hints--

Sie sollten eine Variable namens `discount` mit einer `PercentageDiscount`-Instanz erstellen.

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

Ihre `PercentageDiscount`-Instanz sollte mit `10` Prozent initialisiert werden.

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

Sie sollten `discount.apply_discount(product)` aufrufen und das Ergebnis ausgeben.

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--

--fcc-editable-region--
```
