---
id: 68e9207ddcba053c5cb24413
title: Schritt 2
challengeType: 20
dashedName: step-2
---

# --description--

In Python sind Typangaben optionale Hinweise, die anderen Entwicklern mitteilen, welchen Datentyp eine Variable oder Funktion voraussichtlich hat. Hier ist ein Beispiel, das die Typangabe `str` zum `name`-Parameter hinzufügt.

```py
def greet(name: str):
    return 'Hello, ' + name
```

Aktualisieren Sie für die Parameter der `__init__`-Methode den `name`-Parameter so, dass er die Typangabe `str` erhält, und den `price`-Parameter so, dass er die Typangabe `float` erhält.

# --hints--

Sie sollten die Typangabe `str` zum `name`-Parameter und die Typangabe `float` zum `price`-Parameter in der `__init__`-Methode der `Product`-Klasse hinzufügen.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
--fcc-editable-region--
    def __init__(self, name, price):
--fcc-editable-region--
        self.name = name
        self.price = price
```
