---
id: 68ef53ef276acfa40c83cb12
title: Schritt 10
challengeType: 20
dashedName: step-10
---

# --description--

Fügen Sie der `bool`-Methode einen Rückgabetyp-Hinweis `is_applicable` hinzu, damit klar wird, dass sie ein `bool` zurückgibt.

# --hints--

Ihre `is_applicable`-Methode sollte einen Rückgabetyp-Hinweis `-> bool` haben.

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_returns("bool")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
--fcc-editable-region--
    def is_applicable(self, product: Product, user_tier: str):
--fcc-editable-region--
        pass

product = Product('Wireless Mouse', 50.0)
print(product)
```
