---
id: 68f119ad905b29ff59d858c8
title: Schritt 5
challengeType: 20
dashedName: step-5
---

# --description--

Genau wie im vorherigen Schritt fügen Sie in Ihrer `__str__`-Methode einen Rückgabetyp-Hinweis von `str` hinzu, da diese Methode eine String-Darstellung des Objekts zurückgibt.

# --hints--

Ihre `__str__`-Methode sollte einen Rückgabetyp-Hinweis von `str` haben.

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
