---
id: 67ef8f695db52583424bd118
title: 设计一个电影评价页面
challengeType: 25
dashedName: design-a-movie-review-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。 

**用户故事：**

1. 你应该有一个 `main` 元素。
2. 在 `main` 元素内部，你应该有一个 `h1` 元素，用于显示电影标题。
3. 在 `h1` 元素下面，你应该有一个 `img` 元素，用于显示电影封面。 你的 `img` 元素应包含描述性 `alt` 文本描述图像。 你可以随意使用此图片：`https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png`。
4. 你应该有一个 `p` 元素，其中包含简短的电影描述。
5. 你应该使用另一个 `p` 元素来显示电影评分。 在其中，你应该按照所列顺序排列这些项目：
   - 一个 `strong` 元素，文本为 `Movie Rating:`。
   -  一个 `span` 元素，其 `aria-hidden` 属性设为 `true`，并且包含评分的可视化表示（例如星星 `⭐⭐⭐⭐⭐⭐⭐⭐⭐☆`）。
   - 在 span 之后，显示一个括号括起来的数字评分值，（例如 `9.2/10`）。
6. 你应该有一个 `h2` 元素，文本为 `Cast Members` 。
7. 你应该有一个 `ul` 元素。
8. 在 `ul` 元素内，你应该有多个 `li` 元素，每个元素包含一个用于演员名字的 `strong` 元素，后面跟着以文本 `as` 开头的对应字符名。（例如，`James Holloway as Ethan Carter`）。

# --hints--

你应该有一个 `main` 元素。

```js
assert.isNotNull(document.querySelector("main"));
```

在 `main` 元素内部，你应该有一个包含电影标题的 `h1` 元素。

```js
const h1 = document.querySelector("main > h1");
assert.isAbove(h1?.innerText.trim().length, 0);
```

在 `h1` 元素下面，你应该有一个 `img` 元素，并且带有描述性的 `alt` 属性。
  
```js
const img = document.querySelector("main > h1 + img");
assert.isNotNull(img);
assert.isNotEmpty(img.getAttribute("alt"));
```

你应该在图片后添加一个 `p` 元素，其中包含电影描述。

```js
const description = document.querySelector("main img + p");
assert.isNotNull(description);
assert.isNotEmpty(description.innerText.trim());
```

你应该有另一个包含评分的 `p` 元素。

```js
const paragraphs = document.querySelectorAll("main p");
assert.isAtLeast(paragraphs.length, 2);
```

在第二个 `p` 元素内，你应该有一个带有文本 `Movie Rating: ` 的 `strong` 元素。

```js
const strongEl = document.querySelector("main p:nth-of-type(2) strong");
assert.match(strongEl?.innerText.trim(), /^Movie Rating:?$/);
```

评分段落中应该有一个 `span` 元素。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.isNotNull(spanEl);
```

评分段落元素中的 `span` 元素，应当包含 10 个星星，可以是实心（⭐）或空心（☆），并且在 `span` 后显示一个括号括起来的数字评分值。

```js  
const spanEl = document.querySelector("main p:nth-of-type(2) span");
const fullRatingText = document.querySelector("main p:nth-of-type(2)")?.textContent.replace(/\s+/g, ' ').trim();

assert.match(spanEl?.innerText.trim(), /^[⭐☆]{10}$/);
assert.match(fullRatingText, /Movie Rating:?\s*[⭐☆]{10}\s*\(\s*\d+(?:\.\d+)?\s*\/\s*10\s*\)$/);
```

评分段落元素中的 `span` 元素应该有一个 `aria-hidden` 属性，设置为 `true`。

```js
const spanEl = document.querySelector("main p:nth-of-type(2) span");
assert.strictEqual(spanEl?.getAttribute("aria-hidden"), "true");
```

你应该有一个 `h2` 元素。

```js
assert.exists(document.querySelector("h2"));
```

你的 `h2` 元素应当位于第二个 `p` 元素之后。

```js
const allParagraphs = document.querySelectorAll("p");
assert.isAtLeast(allParagraphs.length, 2);

const h2 = document.querySelector("h2");
const main = document.querySelector("main");

assert.exists(h2);
assert.exists(main);

const elements = Array.from(main.querySelectorAll("*"));
const lastPIndex = elements.indexOf(allParagraphs[1]);
const h2Index = elements.indexOf(h2);

assert.isAbove(h2Index, lastPIndex);
```

你的 `h2` 元素的文本应为 `Cast Members`。

```js
const h2 = document.querySelector("h2");
assert.equal(h2?.textContent.trim(), "Cast Members");
```

你应该在 `h2` 之后有一个 `ul` 元素。

```js
const ul = document.querySelector("main h2 + ul");
assert.isNotNull(ul);
```

`ul` 元素应包含多个 `li` 元素。

```js
const liItems = document.querySelectorAll("main ul li");
assert.isAbove(liItems.length, 1);
```

每个 `li` 应包含一个用于演员名字的 `strong` 元素，后跟以文本 `as` 开头的对应字符名。（例如，`James Holloway as Ethan Carter`）。

```js
const liItems = document.querySelectorAll("main ul li");  
assert.isNotEmpty(liItems);
liItems.forEach(li => {
  const strong = li.querySelector("strong");
  assert.isNotNull(strong);
  const fullText = li.textContent.trim();
  const actorName = strong.textContent.trim();
  assert.isTrue(fullText.startsWith(actorName));
  const characterPart = fullText.slice(actorName.length).trim();
  assert.match(characterPart, /^as\s.+/);
});
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
 <html lang="en">
 <head>
     <meta charset="UTF-8">
     <meta name="viewport" content="width=device-width, initial-scale=1.0">
     <title>Movie Review</title>
 </head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Movie Review</title>
  </head>
  <body>
    <main class="review">
      <h1>Movie Review: Rise Beyond</h1>
      <img
        src="https://cdn.freecodecamp.org/curriculum/labs/rise-beyond-2.png"
        alt="The Rise Beyond movie poster, starring Ethan Carter. A portrait of a young boy standing on a stage with his back toward the audience, holding a guitar at his side, staring at a backdrop of a cityscape with the sun rising behind its tall buildings."
      />
      <p>
        Rise Beyond is an uplifting drama that tells the story of Ethan Carter,
        a young man from a small town who dreams of becoming a world-class
        musician. Facing financial struggles and self-doubt, Ethan embarks on a
        journey of resilience, guided by his mentor Professor Adams and
        childhood friend Lena Mitchell. Through setbacks, failures, and moments
        of triumph, Rise Beyond reminds us that our greatest potential lies
        beyond our fears.
      </p>
      <p>
        <strong>Movie Rating:</strong>
        <span aria-hidden="true">⭐⭐⭐⭐⭐⭐⭐⭐⭐☆</span> (9.2/10)
      </p>

      <h2>Cast Members</h2>
      <ul>
        <li><strong>James Holloway</strong> as Ethan Carter</li>
        <li><strong>Olivia Sterling</strong> as Lena Mitchell</li>
        <li><strong>William Lancaster</strong> as Professor Adams</li>
        <li>
          <strong>Maria Collins</strong> as Evelyn Carter (Ethan’s Mother)
        </li>
      </ul>
    </main>
  </body>
</html>
```
