---
id: 68cabc534d2e741f33a3fe1d
title: Construye un Rastreador de Estadísticas de Personajes de Juego
challengeType: 27
dashedName: lab-game-character-stats
---

# --description--

En este laboratorio, construirás un rastreador de estadísticas de personajes de juego. El programa te permitirá crear un personaje con atributos específicos, actualizar esos atributos y obtener las estadísticas actuales del personaje.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Crea una clase llamada `GameCharacter` que representa un personaje del juego y gestiona las estadísticas del personaje.
2. When instantiated, a new `GameCharacter` object should have the following attributes:
    - `_name` establecido con la cadena dada en el momento de la instanciación.
    - `_health` establecido en `100`.
    - `_mana` establecido en `50`.
    - `_level` establecido en `1`.
3. Crea una propiedad `name` para acceso de solo lectura al nombre del personaje.
4. For the `health` property:
     - Define un getter que devuelve la salud actual.
     - Define un setter que impida que la salud se establezca por debajo de `0` y limite la salud a `100`.
5. For the `mana` property:
     - Define un getter que devuelva el mana actual.
     - Define un setter que impida que el mana se establezca por debajo de `0` y limite el mana a `50`.
6. Crea un getter para `level` que devuelva el nivel actual del personaje.
7. Define a method named `level_up` that:
    - Aumenta el nivel del personaje en 1.
    - Restablece la salud a `100` y el mana a `50` usando sus correspondientes setters de propiedad.
    - Imprime un mensaje en la forma `<name> leveled up to <level>!` (donde `<name>` y `<level>` deben ser reemplazados por el nombre del personaje y el nuevo nivel, respectivamente).
8. Define a `__str__` method that returns a formatted string including:
    - El nombre del personaje.
    - El nivel del personaje.
    - La salud actual del personaje.
    - El mana actual del personaje.  
Por ejemplo, un personaje llamado `Kratos`, justo después de la instanciación, debería representarse como lo siguiente:

   ```md
   Name: Kratos
   Level: 1
   Health: 100
   Mana: 50
   ```

## Ejemplo de uso

```python
hero = GameCharacter('Kratos') # Creates a new character named Kratos
print(hero)  # Displays the character's stats

hero.health -= 30  # Decreases health by 30
hero.mana -= 10    # Decreases mana by 10
print(hero)  # Displays the updated stats

hero.level_up()  # Levels up the character
print(hero)  # Displays the stats after leveling up
```

# --hints--

Debes tener una clase `GameCharacter`.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("GameCharacter")`)) })
```

Debes tener un método `__init__` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__init__")`) 
})
```

El método `__init__` debe tener dos parámetros, siendo el primero `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.__init__).parameters)
    assert params[0] == "self"
    assert len(params) == 2
    `) 
})
```

Debes asignar `self._name` a la cadena dada en el momento de la instanciación dentro del método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._name == "Kratos"
    `) 
})
```

Debes asignar `100` a `self._health` dentro del método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._health == 100
    `) 
})
```

Debes asignar `50` a `self._mana` dentro del método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._mana == 50
    `) 
})
```

Debes asignar `1` a `self._level` dentro del método `__init__`.

```js
({ 
    test: () => runPython(`
    character = GameCharacter("Kratos")
    assert character._level == 1
    `) 
})
```

Debes tener un método `name` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("name")`) 
})
```

El método `name` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_args("self")`) 
})
```

Debes retornar `self._name` en el método `name`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_return("self._name")`)
})
```

El método `name` debe tener un decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("name").has_decorators("property")`) 
})
```

Debes tener un método `health` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("health")`) 
})
```

El método `health` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_args("self")`) 
})
```

Debes retornar `self._health` en el método `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_return("self._health")`)
})
```

El método `health` debe tener un decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("health").has_decorators("property")`) 
})
```

Debes crear un método setter `health`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1]`) 
})
```

El método setter de `health` debe tener un decorador `@health.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("health")[1].has_decorators("health.setter")`) 
})
```

El método setter `health` debe tener dos parámetros, siendo el primero `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.health.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Debes establecer `self._health` en `0` si el valor dado al setter `health` es menor que `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = -1
    assert char.health == 0
    char.health = -10
    assert char.health == 0
`)})
```

Debes establecer `self._health` en `100` si el valor dado al setter `health` es mayor que `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 101
    assert char.health == 100
    char.health = 200
    assert char.health == 100
`)})
```

Debes establecer `self._health` al valor dado al setter `health` cuando el valor esté entre `0` y `100`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.health = 95
    assert char.health == 95
    char.health = 1
    assert char.health == 1
`)})
```

Deberías tener un método `mana` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("mana")`) 
})
```

El método `mana` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_args("self")`) 
})
```

Debes retornar `self._mana` en el método `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_return("self._mana")`)
})
```

El método `mana` debe tener un decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("mana").has_decorators("property")`) 
})
```

Debes crear un método setter `mana`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1]`) 
})
```

El método setter de `mana` debe tener un decorador `@mana.setter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_functions("mana")[1].has_decorators("mana.setter")`) 
})
```

El método setter `mana` debe tener dos parámetros, siendo el primero `self`.

```js
({ 
    test: () => runPython(`
    import inspect
    params = list(inspect.signature(GameCharacter.mana.fset).parameters)
    assert params[0] == "self"
    assert len(params) == 2`) 
})
```

Debes establecer `self._mana` en `0` si el valor dado al setter `mana` es menor que `0`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = -1
    assert char.mana == 0
    char.mana = -10
    assert char.mana == 0
`)})
```

Debes establecer `self._mana` en `50` si el valor dado al setter `mana` es mayor que `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 51
    assert char.mana == 50
    char.mana = 100
    assert char.mana == 50
`)})
```

Debes asignar `self._mana` al valor dado al setter `mana` si el valor está entre `0` y `50`.

```js
({ 
    test: () => runPython(`
    char = GameCharacter("test")
    char.mana = 25
    assert char.mana == 25
    char.mana = 0
    assert char.mana == 0
    char.mana = 50
    assert char.mana == 50
`)})
```

Deberías tener un método `level` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level")`) 
})
```

El método `level` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_args("self")`) 
})
```

Debes retornar `self._level` en el método `level`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_return("self._level")`)
})
```

El método `level` debe tener un decorador `@property`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level").has_decorators("property")`) 
})
```

Deberías tener un método `level_up` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("level_up")`) 
})
```

El método `level_up` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_args("self")`) 
})
```

Debes aumentar `self._level` en `1` en el método `level_up`.

```js
({ 
  test: () => runPython(`
    node = _Node(_code).find_class("GameCharacter").find_function("level_up")
    assert (
        node.has_stmt("self._level += 1") or
        node.has_stmt("self._level = self._level + 1")
    )
`)
})
```

Debes establecer `self.health` en `100` en el método `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.health = 100")`) 
})
```

Debes establecer `self.mana` en `50` en el método `level_up`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("level_up").has_stmt("self.mana = 50")`) 
})
```

El método `level_up` debería imprimir `<name> leveled up to <level>!` (donde `<name>` y `<level>` deben ser reemplazados por el nombre del personaje y el nuevo nivel, respectivamente).

```js
({
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.level_up()
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
            
    assert "test leveled up to 2!" in output
`)
})
```

Debes tener un método `__str__` en la clase `GameCharacter`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").has_function("__str__")`) 
})
```

El método `__str__` debe tener solo un parámetro, `self`.

```js
({ 
    test: () => runPython(`assert _Node(_code).find_class("GameCharacter").find_function("__str__").has_args("self")`) 
})
```

Tu método `__str__` debe devolver una cadena con las estadísticas del personaje usando el formato proporcionado.

```js
({ 
  test: () => runPython(`
    import io
    import sys
            
    captured_output = io.StringIO()
    sys.stdout = captured_output
            
    gc = GameCharacter("test")
    gc.mana = 15
    gc.health = 77
    print(gc)
            
    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()
    expected = """Name: test
    Level: 1
    Health: 77
    Mana: 15"""
    assert expected in output
`)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
class GameCharacter:
    def __init__(self, name):
        self._name = name
        self._health = 100
        self._mana = 50
        self._level = 1

    @property
    def name(self):
        return self._name

    @property
    def health(self):
        return self._health

    @health.setter
    def health(self, value):
        if value < 0:
            self._health = 0
        elif value > 100:
            self._health = 100
        else:
            self._health = value

    @property
    def mana(self):
        return self._mana

    @mana.setter
    def mana(self, value):
        if value < 0:
            self._mana = 0
        elif value > 50:
            self._mana = 50
        else:
            self._mana = value

    @property
    def level(self):
        return self._level

    def level_up(self):
        self._level += 1
        self.health = 100
        self.mana = 50
        print(f"{self._name} leveled up to {self._level}!")

    def __str__(self):
        return (
            f"Name: {self._name}\n"
            f"Level: {self._level}\n"
            f"Health: {self._health}\n"
            f"Mana: {self._mana}"
        )
```
