---
id: 65c4f195de7e2ad5932be717
title: Крок 9
challengeType: 20
dashedName: step-9
---

# --description--

Далі потрібно визначити механізм для вставлення вузлів у дерево. Для цього потрібно визначити метод `_insert`, який є допоміжною функцією і буде використовуватися пізніше справжнім методом `insert`.

Цей метод є рекурсивним, тобто він викликає сам себе, щоб пройти дерево, доки не знайде відповідне місце для нового вузла.

Визначте метод `_insert` з параметрами `self`, `node` та `key`.

# --hints--

Ви повинні визначити метод `_insert` всередині класу `BinarySearchTree`. Не забувайте використовувати `pass`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\([^(]*\)\s*:/m) })
```

Ваш метод `_insert` має приймати три параметри: `self`, `node` та `key`.

```js
({ test: () => assert.match(code, /def\s+_insert\s*\(\s*self\s*,\s*node\s*,\s*key\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--
```
