---
id: 65c4f300da28d8da361bfa93
title: Крок 14
challengeType: 20
dashedName: step-14
---

# --description--

Усередині вашої умови `elif` замініть `pass` викликом методу `_insert` з правим дочірнім елементом поточного вузла як першим аргументом і `key` як другим аргументом. Призначте результат атрибуту `right` поточного вузла.

# --hints--

Ви повинні видалити ключове слово `pass` з блоку `elif`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Ви повинні викликати метод `self._insert()`, передаючи `node.right` і `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```

Ви повинні призначити результат вашого виклику `self._insert()` правильному атрибуту (`node.right`) поточного вузла.

```js
({ test: () => assert.match(code, /node\.right\s*=\s*self\._insert\(\s*node\.right\s*,\s*key\s*\)/) });
```


# --seed--

## --seed-contents--


```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

--fcc-editable-region--
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            pass

--fcc-editable-region--

```
