---
id: 65c4ef8463f869d0eaf87c67
title: Крок 2
challengeType: 20
dashedName: step-2
---

# --description--

Всередині класу `TreeNode` замініть `pass` на метод `__init__`, щоб ініціалізувати атрибути об’єкта. Поки що не додавайте жодних параметрів.

# --hints--

Ви повинні видалити ключове слово `pass` з класу `TreeNode` і перемістити його всередину методу `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

Ви повинні визначити метод `__init__` всередині вашого класу `TreeNode`. Не забувайте використовувати ключове слово `pass` всередині.

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
