---
id: 65c4f07da204f4d2f325fbd3
title: Крок 5
challengeType: 20
dashedName: step-5
---

# --description--

У методі `__init__` ініціалізуйте атрибути `left` та `right` вузла значенням `None`. Це потрібно, бо коли вузол створюється вперше, він не має дочірніх елементів зліва чи справа. Не забудьте використати ключове слово `self`.

# --hints--

Ви повинні призначити `None` атрибуту `left` вузла, використовуючи `self.left`.

```js
({ test: () => assert.match(code, /self\.left\s*=\s*None/m) })
```

Ви повинні призначити `None` атрибуту `right` вузла, використовуючи `self.right`.

```js
({ test: () => assert.match(code, /self\.right\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:    
    def __init__(self, key):
        self.key = key
--fcc-editable-region--
```
