---
id: 65c4f3258d2e4cdacc919dfd
title: Крок 15
challengeType: 20
dashedName: step-15
---

# --description--

Наприкінці вашого методу `_insert`, після завершення процесу вставки, поверніть поточний вузол, щоб оновити структуру дерева на вищих рівнях рекурсивного стека викликів.


# --hints--

Ви повинні повернути поточний вузол поза межами умовних блоків.

```js
({ test: () => assert.match(code, /^\s{8}return\s+node/m) })
```

# --seed--

## --seed-contents--


```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

class BinarySearchTree:
    def __init__(self):
        self.root = None

--fcc-editable-region--    
    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)
        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:
            node.right = self._insert(node.right, key)

--fcc-editable-region--

```
