---
id: 65c4eb814cc977c95cd7df0e
title: Крок 1
challengeType: 20
dashedName: step-1
---

# --description--

У цьому проєкті ви створите двійкове дерево пошуку (BST). BST — це структура даних, у якій кожен вузол має не більше двох дочірніх елементів: лівий дочірній елемент містить значення, менші за батьківський вузол, а правий — значення, більші за батьківський вузол, що дозволяє ефективно виконувати операції пошуку та сортування.

Ось як виглядає двійкове дерево пошуку:

<img alt="a binary search tree" src="https://cdn.freecodecamp.org/curriculum/python/bst-example.png" style="background-color: white; height:500px; width:500px; padding: 10px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

Почніть із оголошення порожнього класу `TreeNode`. Клас `TreeNode` представляє вузол у двійковому дереві пошуку. Використайте ключове слово `pass`, щоб заповнити тіло класу і уникнути помилки.

# --hints--

Ви повинні використати ключове слово `class`, щоб оголосити порожній клас із назвою `TreeNode`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(TreeNode)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
