---
id: 65c4f0d32ca17ad4b1636b0e
title: Крок 8
challengeType: 20
dashedName: step-8
---

# --description--

У методі `__init__` видаліть `pass` і ініціалізуйте `root` значенням `None`.

Атрибут `root` представляє кореневий вузол двійкового дерева пошуку. Оскільки це конструктор, коли створюється новий об’єкт `BinarySearchTree`, він починається з порожнього дерева, тому атрибут `root` встановлюється в `None`.

# --hints--

Ви повинні видалити інструкцію `pass` з методу `__init__`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});

```

Ви повинні ініціалізувати атрибут `root` значенням `None` за допомогою `self.root`.

```js
({ test: () => assert.match(code, /^\s{8}self\.root\s*=\s*None/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None
--fcc-editable-region--
class BinarySearchTree:
    def __init__(self):
        pass
--fcc-editable-region--        
```
