---
id: 65c4f02dbd3557d23f12db62
title: Крок 4
challengeType: 20
dashedName: step-4
---

# --description--

Всередині методу `__init__` видаліть `pass` і призначте значення параметра `key` атрибуту `key` вузла, використовуючи `self.key`.

Це означає, що атрибут `key` екземпляра `TreeNode` буде призначений значенням, переданим під час створення об’єкта.

# --hints--

Ви повинні видалити інструкцію `pass` з методу `__init__`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("TreeNode").find_function("__init__").has_pass()`
      )
    );
  },
});

```

Ви повинні призначити значення параметра `key` атрибуту `key` вузла, використовуючи `self.key`.

```js
({ test: () => assert.match(code, /^\s{8}self\.key\s*=\s*key/m) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    def __init__(self, key):
        pass
--fcc-editable-region--
```
