---
id: 65c4f22498d22ed775ef8efb
title: Крок 10
challengeType: 20
dashedName: step-10
---

# --description--

Тепер потрібно перевірити, чи параметр `node` є `None`. Якщо так, це означає, що метод досяг листового вузла або порожнього місця в дереві, куди слід вставити новий вузол.

У тілі методу `_insert` замініть `pass` на інструкцію `if`, яка перевіряє, чи `node is None`.

У новому блоці `if` поверніть `TreeNode(key)`, щоб створити новий екземпляр `TreeNode` з наданим ключем. Це стане новим листовим вузлом, фактично вставляючи ключ у дерево.

# --hints--

У вашому методі `pass` не повинно бути `_insert`.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("BinarySearchTree").find_function("_insert").has_pass()`)) })

```

Ви повинні написати інструкцію `if`, щоб перевірити, чи `node is None`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^\s{8}if\s+node\s+is\s+None\s*:/m));
  }
})
```

Ви повинні повернути `TreeNode(key)` з блоку `if`.

```js
({ test: () =>
  {
    const transformedCode = code.replace(/\r/g, "");        
    const ins = __helpers.python.getDef("\n"+transformedCode, "_insert");
    const {function_body} = ins;    
    assert(function_body.match(/^(\s{8})if\s+node\s+is\s+None\s*:\s*^\1\s{4}return\s+TreeNode\s*\(\s*key\s*\)/m));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        pass
--fcc-editable-region--
```
