---
id: 65c4f2b7178afed8e88f782f
title: Крок 12
challengeType: 20
dashedName: step-12
---

# --description--

Якщо `key` < `node.key` оцінюється як `True`, тоді новий вузол слід розмістити в лівому піддереві.

Видаліть `pass` і рекурсивно викличте метод `_insert`, передаючи лівий дочірній елемент як перший аргумент і `key` як другий аргумент. Призначте результат атрибуту `left` поточного вузла.

# --hints--

Ви повинні видалити ключове слово `pass` з блоку `if`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("_insert").find_if("key < node.key").find_bodies()[0].has_pass()`
      )
    );
  },
});
```

Ви повинні викликати метод `self._insert`, передаючи `node.left` і `key` як аргументи.

```js
({ test: () => assert.match(code, /self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

Ви повинні призначити результат виклику `self._insert()` атрибуту `left` поточного вузла.

```js
({ test: () => assert.match(code, /node\.left\s*=\s*self\._insert\(\s*node\.left\s*,\s*key\s*\)/) });
```

# --seed--

## --seed-contents--

```py
class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None
--fcc-editable-region--    

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            pass
--fcc-editable-region--    

```
