---
id: 65c4f09e074dd8d37830ea00
title: Крок 6
challengeType: 20
dashedName: step-6
---

# --description--

Створіть ще один порожній клас із назвою `BinarySearchTree`, який представляє двійкове дерево пошуку.

# --hints--

У вас має бути клас із назвою `BinarySearchTree`. Пам’ятайте, що потрібно використовувати ключове слово `pass`.

```js
({
  test: () => {
    assert(runPython(`
      import inspect
      inspect.isclass(BinarySearchTree)
    `));
  }
})
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--
```
