---
id: 65c63e6962159fa372ecdbec
title: Крок 22
challengeType: 20
dashedName: step-22
---

# --description--

Якщо друга інструкція `if` не є `True`, це означає, що цільовий ключ більший або дорівнює ключу поточного вузла.

У двійковому дереві пошуку, якщо цільовий ключ більший за ключ поточного вузла, пошук продовжується у правому піддереві.

Після блоку `if` поверніть результат виклику методу `_search` з правим дочірнім елементом поточного вузла та ключем як аргументами.

# --hints--

Ви повинні рекурсивно викликати метод `_search`, використовуючи `self._search()`.

```js
const after_split = code.split('return self._search(node.left, key)')[1];
assert.match(after_split, /self\._search/);
```

Ви повинні передати два аргументи у виклик методу `_search`: `node.right` та `key`.

```js
({ test: () => assert.match(code, /self\._search\(\s*node\.right\s*,\s*key\s*\)/) });
```

Ви повинні повернути результат виклику методу `_search`.

```js
({
  test: () => assert.match(code, /return\s+self\._search\(\s*node\.right\s*,\s*key\s*\)/)
});
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)

--fcc-editable-region--
```
