---
id: 65c63fd9b387f0a6c5bf6a72
title: Крок 24
challengeType: 20
dashedName: step-24
---

# --description--

У методі `search` видаліть `pass` і викличте допоміжний метод `_search` з такими аргументами.

- `self.root`: Це корінь двійкового дерева пошуку. Пошук починається з кореня.
- `key`: Це значення, яке користувач хоче знайти у двійковому дереві пошуку.

Внутрішньо `search` делегує фактичну логіку пошуку допоміжному методу `_search`, який виконує рекурсивний пошук у двійковому дереві пошуку.

# --hints--

Ви повинні видалити ключове слово `pass` з методу `search`.

```js
({
  test: () => {
    assert.isFalse(
      runPython(
        `_Node(_code).find_class("BinarySearchTree").find_function("search").has_pass()`
      )
    );
  },
});

```

Ви повинні викликати метод `_search` всередині методу `search`.

```js
({ test: () => assert(runPython(`
  str(_Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0]).startswith("self._search(")
`)) });
```

Ви повинні викликати метод `_search`, передаючи `self.root` і `key` як аргументи.

```js
({ test: () => assert(runPython(`
  _Node(_code).find_class("BinarySearchTree").find_function("search").find_calls("_search")[0].is_equivalent("self._search(self.root, key)")
`)) });
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
--fcc-editable-region--
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        pass
--fcc-editable-region--
```
