---
id: 65c6409418806da8d0636ffc
title: Крок 26
challengeType: 20
dashedName: step-26
---

# --description--

Функціональність `insert` і `search` завершена, настав час їх протестувати.

Зверніть увагу, що на цьому етапі вузли не відсортовані, а просто вставлені. Пізніше ви попрацюєте над сортуванням за допомогою обходу inorder.

Ви можете створити екземпляр класу в Python ось так:

```python
object_name = ClassName()
```

Поза межами визначень класів створіть екземпляр класу `BinarySearchTree` і призначте його змінній `bst`.


# --hints--

Ви повинні створити екземпляр класу `BinarySearchTree`.

```js
assert.match(code, /BinarySearchTree\(\s*\)/);
```

Ви повинні призначити новий екземпляр `BinarySearchTree` змінній `bst`.

```js
assert.match(code, /bst\s*=\s*BinarySearchTree\(\s*\)/);
``` 

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None


class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--


--fcc-editable-region--
```
