---
id: 65c646d4148ae3b2d1cbcac4
title: Крок 30
challengeType: 20
dashedName: step-30
---

# --description--

Зверніть увагу, що ваш пошук повертає щось на кшталт `80: <__main__.TreeNode object at 0x108b3e0>`. Це рядкове представлення за замовчуванням при виведенні екземпляра класу.

Щоб змінити це і виводити корисне значення, визначте інший метод з назвою `__str__` у класі `TreeNode`. Він приймає один аргумент `self`.

Після визначення `__str__` ви отримаєте виняток у консолі, бо метод `__str__` поки що нічого не повертає. Ви працюватимете над тілом методу на наступному кроці.

# --hints--

Ви маєте визначити метод `__str__`, який приймає один аргумент `self`. Не забувайте використовувати `pass`.

```js
assert.match(code, /^(\s+)def\s+__init__.+?^\1def\s+__str__\(\s*self\s*\)\s*:\s*\n^\1\1pass/ms)
```


# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--

--fcc-editable-region--

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

print('Search for 80:', bst.search(80))

```
