---
id: 65ca075e3184180309f4c1d9
title: Крок 39
challengeType: 20
dashedName: step-39
---

# --description--

Всередині блоку `else` замініть `pass` на інструкцію `if`, щоб перевірити, чи є лівий дочірній елемент поточного вузла `None`.

Коли `node.left` дорівнює `None`, лівого дочірнього елемента немає. Тому поверніть правий дочірній елемент із нового блоку `if` як заміну.

# --hints--

Ви повинні видалити ключове слово `pass` із блоку `else`.

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree").find_function("_delete"))`
    );
    assert.notInclude(pyClassStr, "pass");
  },
});
```

Ви повинні написати інструкцію `if`, яка перевіряє, чи `node.left` дорівнює `None`.

```js
assert.match(code, /if\s+node\.left\s*is\s*None/);
```

Ви повинні повернути правий дочірній елемент поточного вузла.

```js 
assert.match(code, /return\s+node\.right/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key) 
        else:
            pass   
--fcc-editable-region--

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
