---
id: 65d8a6fcb15a3a239ba35dfd
title: Крок 46
challengeType: 20
dashedName: step-46
---

# --description--

Після того, як ви знайдете мінімальне значення, потрібно рекурсивно видалити вузол із мінімальним значенням із правого піддерева.

Цей крок гарантує, що вузол із мінімальним значенням буде видалений із дерева, при цьому зберігається властивість двійкового дерева пошуку (BST).

Рекурсивно викличте метод `_delete` з аргументами `node.right` і `node.key`. Призначте повернене значення виклику `_delete()` правому дочірньому елементу поточного вузла.

# --hints--

Ви повинні рекурсивно викликати метод `_delete` з аргументами `node.right` і `node.key`.

```js
assert.match(code, /self\._delete\(\s*node\.right\s*,\s*node\.key\s*/);
```

Ви повинні призначити повернене значення виклику `_delete()` правому дочірньому елементу поточного вузла.

```js
assert.match(code, /node\.right\s*=\s*self\._delete\(\s*node\.right\s*,\s*node\.key/);
```

# --seed--

## --seed-contents--

```py

class TreeNode:

    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

    def __str__(self):
        return str(self.key)

class BinarySearchTree:

    def __init__(self):
        self.root = None

    def _insert(self, node, key):
        if node is None:
            return TreeNode(key)

        if key < node.key:
            node.left = self._insert(node.left, key)
        elif key > node.key:

            node.right = self._insert(node.right, key)
        return node

    def insert(self, key):
        self.root = self._insert(self.root, key)
        
    def _search(self, node, key):
        if node is None or node.key == key:
            return node
        if key < node.key:
            return self._search(node.left, key)
        return self._search(node.right, key)
    
    def search(self, key):
        return self._search(self.root, key)

--fcc-editable-region--
    def _delete(self, node, key):
        if node is None:
            return node
        if key < node.key:
            node.left = self._delete(node.left, key)
        elif key > node.key:
            node.right = self._delete(node.right, key)
        else:
            if node.left is None:
                return node.right
            elif node.right is None:
                return node.left

            node.key = self._min_value(node.right) 
            
--fcc-editable-region--
    def _min_value(self, node):
        while node.left is not None:
            node = node.left

bst = BinarySearchTree()

nodes = [50, 30, 20, 40, 70, 60, 80]

for node in nodes:
    bst.insert(node)

# print('Search for 80:', bst.search(80))


```
