---
id: 695774002591bbc5f8cf3e53
title: Створіть функцію застосування знижки
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

У цій лабораторній роботі ви напишете функцію, яка обчислює фінальну ціну елемента після застосування відсоткової знижки.

Наприклад, якщо ціна елемента дорівнює `50`, а застосовується знижка `20`, сума знижки становить `10`, а фінальна ціна — `40`.

**Мета:** Виконайте наведені нижче історії користувача, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні визначити функцію з іменем `apply_discount`.

2. Функція `apply_discount` має приймати рівно два параметри: `price` і `discount`.

3. Якщо `price` не є числом (`int` або `float`), функція має повертати рядок `The price should be a number`.

4. Якщо `discount` не є числом (`int` або `float`), функція має повертати рядок `The discount should be a number`.

5. Якщо `price` менше або дорівнює `0`, функція має повертати рядок `The price should be greater than 0`.

6. Якщо `discount` менше за `0` або більше за `100`, функція має повертати рядок `The discount should be between 0 and 100`.

7. Якщо обидва вхідні дані є дійсними, функція має обчислити знижку як відсоток від ціни.

8. Функція має повертати фінальну ціну після застосування знижки.

# --hints--

Ви повинні мати функцію з іменем `apply_discount`.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Ваша функція `apply_discount` має приймати два параметри: `price` і `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Коли `apply_discount` викликається з `price` (першим аргументом), який не є числом (`int` або `float`), вона має повертати `The price should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Коли `apply_discount` викликається з `discount` (другим аргументом), який не є числом (`int` або `float`), вона має повертати `The discount should be a number`.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Коли `apply_discount` викликається з `price`, що менше або дорівнює `0`, вона має повертати `The price should be greater than 0`.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Коли `apply_discount` викликається з `discount`, що менше за `0` або більше за `100`, вона має повертати `The discount should be between 0 and 100`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` має повертати `80`.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` має повертати `100`.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` має повертати `50`.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Коли `apply_discount` викликається зі знижкою `100`, вона має повертати `0`.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` має повертати `59.6`.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
