---
id: 684aaf9ec670c68d20efd0d0
title: Construye un Gestor de Configuración de Usuario
challengeType: 27
dashedName: build-a-user-configuration-manager
---

# --description--

En este laboratorio, construirás un Gestor de Configuración de Usuario que permite a los usuarios administrar sus ajustes como tema, idioma y notificaciones. Implementarás funciones para agregar, actualizar, eliminar y ver los ajustes de usuario.

**Objetivo:** Cumplir con las historias de usuario a continuación y pasar todas las pruebas para completar el laboratorio.

**Historias de usuario:**

1. Debes definir una función llamada `add_setting` con dos parámetros que representen un diccionario de configuraciones y una tupla que contenga un par clave-valor

1. `add_setting` function should:
    
   - Convierte la key y value a minúsculas.
   - Si la configuración clave existe, devuelve `Setting '[key]' already exists! Cannot add a new setting with this name.`
   - Si la configuración de la clave no existe, agrega el par clave-valor al diccionario de configuraciones dado y devuelve `Setting '[key]' added with value '[value]' successfully!`.
   - Los mensajes devueltos deben tener la clave y el valor en minúsculas.

1. Debes definir una función llamada `update_setting` con dos parámetros que representan un diccionario de configuraciones y una tupla que contiene un par clave-valor.

1. `update_setting` function should:
       
   - Convierte la key y el value a minúsculas.
   - Si la configuración de la clave existe, actualiza su valor en el diccionario de configuraciones dado y devuelve: `Setting '[key]' updated to '[value]' successfully!`
   - Si la configuración de la clave no existe, devuelve `Setting '[key]' does not exist! Cannot update a non-existing setting.`
   - Los mensajes devueltos deben tener la clave y el valor en minúsculas.


1. Debes definir una función llamada `delete_setting` con dos parámetros que representan un diccionario de configuraciones y una clave.

1. `delete_setting` function should:
    
   - Convierte la key pasada a minúsculas.
   - Si la configuración de la clave existe, elimina el par clave-valor del diccionario de configuraciones dado y devuelve `Setting '[key]' deleted successfully!`
   - Si la configuración de la clave no existe, devuelve `Setting not found!`
   - Los mensajes devueltos deben tener la clave en minúsculas.

1. Debes definir una función llamada `view_settings` con un parámetro que represente un diccionario de configuraciones.

1. `view_settings` function should:
   
   - Devuelve `No settings available.` si el diccionario dado de configuraciones está vacío.
   - Si el diccionario contiene alguna configuración, devuelve una cadena que muestre las configuraciones. La cadena debe comenzar con `Current User Settings:` seguida de los pares clave-valor, cada uno en una nueva línea y con la clave en mayúscula. Por ejemplo, `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})` debería devolver:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. Para probar el código, debes crear un diccionario llamado `test_settings` para almacenar algunas preferencias de configuración del usuario.

# --hints--

Debes crear un diccionario llamado `test_settings` y agregar algunos valores a él.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

Debes definir una función llamada `add_setting`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

La función `add_setting` debe tener dos parámetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting` debería convertir la clave a minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting` debería convertir el valor a minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))` debería devolver el mensaje de error `Setting 'theme' already exists! Cannot add a new setting with this name.`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))` debería agregar un nuevo par clave-valor y devolver el mensaje de éxito `Setting 'volume' added with value 'high' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting` debería agregar correctamente el par clave-valor dado al diccionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

Debes definir una función llamada `update_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

La función `update_setting` debe tener dos parámetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

La función `update_setting` debe convertir la clave a minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

La función `update_setting` debe convertir value a minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))` debería actualizar una clave existente y devolver el mensaje de éxito `Setting 'theme' updated to 'dark' successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))` debería devolver el mensaje de error `Setting 'volume' does not exist! Cannot update a non-existing setting.` cuando la clave no exista.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting` debería actualizar correctamente el par clave-valor dado en el diccionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

Debes definir una función llamada `delete_setting`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

La función `delete_setting` debe tener dos parámetros.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting` debería convertir la clave a minúsculas.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')` debería eliminar la clave existente y devolver el mensaje de éxito `Setting 'theme' deleted successfully!`.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting` debería devolver el mensaje de error `Setting not found!` cuando la clave no exista.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting` debería eliminar correctamente la clave dada del diccionario.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

Debes definir una función llamada `view_settings`.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

La función `view_settings` debe tener un parámetro.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings` debería devolver el mensaje `No settings available.` si el diccionario dado está vacío.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings` debe devolver configuraciones formateadas para un diccionario no vacío.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings` debe capitalizar la primera letra de cada nombre de configuración.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings` debería mostrar los resultados correctos y terminar con un carácter de nueva línea.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: English\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "English",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
