---
id: 5a24c314108439a4d4036151
title: استخدام جملة التبديل للتعامل مع عدة إجراءات
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

يمكنك إخبار مخزن Redux بكيفية التعامل مع عدة أنواع من الإجراءات. لنفترض أنك تدير توثيق المستخدم في مخزن Redux. تريد أن يكون لديك تمثيل للحالة عندما يكون المستخدمون مسجلين دخول وعندما يكونون مسجلين خروج. تمثل هذا بكائن حالة واحد يحتوي على الخاصية `authenticated`. تحتاج أيضًا إلى منشئي إجراءات ينشئون إجراءات تتوافق مع تسجيل دخول المستخدم وتسجيل خروجه، إلى جانب كائنات الإجراءات نفسها.

# --instructions--

محرر الشفرة يحتوي على مخزن وإجراءات ومنشئي إجراءات معدة لك. املأ دالة `reducer` للتعامل مع عدة إجراءات توثيق. استخدم جملة `switch` في JavaScript داخل `reducer` للرد على أحداث الإجراءات المختلفة. هذا نمط قياسي في كتابة مخفضات Redux. يجب أن تقوم جملة التبديل بالتبديل على `action.type` وإرجاع حالة التوثيق المناسبة.

**ملاحظة:** في هذه المرحلة، لا تقلق بشأن عدم تغير الحالة، لأنها صغيرة وبسيطة في هذا المثال. لكل إجراء، يمكنك إرجاع كائن جديد — مثلًا `{authenticated: true}`. ولا تنسَ كتابة حالة `default` في جملة التبديل التي تُرجع الحالة الحالية `state`. هذا مهم لأن تطبيقك قد يحتوي على عدة مخفضات، وكلها تُنفذ عند إرسال أي إجراء، حتى لو لم يكن الإجراء متعلقًا بذلك المخفض. في هذه الحالة، تريد التأكد من إرجاع الحالة الحالية `state`.

# --hints--

استدعاء الدالة `loginUser` يجب أن يُرجع كائنًا يحتوي على خاصية type مضبوطة على السلسلة النصية `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

استدعاء الدالة `logoutUser` يجب أن يُرجع كائنًا يحتوي على خاصية type مضبوطة على السلسلة النصية `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

يجب تهيئة المخزن بكائن يحتوي على خاصية `authenticated` مضبوطة على `false`.

```js
assert(store.getState().authenticated === false);
```

إرسال إجراء `loginUser` يجب أن يُحدث خاصية `authenticated` في حالة المخزن إلى `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

إرسال إجراء `logoutUser` يجب أن يُحدث خاصية `authenticated` في حالة المخزن إلى `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

يجب أن تتعامل دالة `authReducer` مع عدة أنواع من الإجراءات باستخدام جملة `switch`.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
