---
id: 5a24c314108439a4d4036152
title: استخدام const لأنواع الإجراءات
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

ممارسة شائعة عند العمل مع Redux هي تعيين أنواع الإجراءات كثوابت للقراءة فقط، ثم الرجوع إلى هذه الثوابت في كل مكان تُستخدم فيه. يمكنك إعادة هيكلة الشفرة التي تعمل عليها لكتابة أنواع الإجراءات كإعلانات `const`.

# --instructions--

أعلن عن `LOGIN` و `LOGOUT` كقيم `const` وعيّن لهما السلاسل النصية `'LOGIN'` و `'LOGOUT'` على التوالي. ثم حرر `authReducer()` ومنشئي الإجراءات ليشيروا إلى هذه الثوابت بدلًا من القيم النصية.

**ملاحظة:** عادةً ما يكون من الاصطلاح كتابة الثوابت بأحرف كبيرة بالكامل، وهذه ممارسة قياسية في Redux أيضًا.

# --hints--

استدعاء الدالة `loginUser` يجب أن يُرجع كائنًا بخاصية `type` مُعيَّنة إلى السلسلة النصية `LOGIN`.

```js
assert(loginUser().type === 'LOGIN');
```

استدعاء الدالة `logoutUser` يجب أن يُرجع كائنًا بخاصية `type` مُعيَّنة إلى السلسلة النصية `LOGOUT`.

```js
assert(logoutUser().type === 'LOGOUT');
```

يجب تهيئة المخزن بكائن بخاصية `login` مُعيَّنة إلى `false`.

```js
assert(store.getState().authenticated === false);
```

إرسال `loginUser` يجب أن يُحدِّث خاصية `login` في حالة المخزن إلى `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

إرسال `logoutUser` يجب أن يُحدِّث خاصية `login` في حالة المخزن إلى `false`.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

يجب أن تتعامل دالة `authReducer` مع أنواع إجراءات متعددة باستخدام جملة التبديل.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

يجب إعلان `LOGIN` و `LOGOUT` كقيم `const` ويجب تعيينهما إلى السلاسل النصية `LOGIN` و `LOGOUT`.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

يجب أن تشير منشئو الإجراءات والمخفض إلى ثوابت `LOGIN` و `LOGOUT`.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
