---
id: 5a24c314108439a4d403614d
title: تعريف إجراء Redux
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

بما أن Redux هو إطار عمل لإدارة الحالة، فإن تحديث الحالة هو أحد مهامه الأساسية. في Redux، تُحفَّز جميع تحديثات الحالة عن طريق إرسال إجراءات. الإجراء هو ببساطة كائن JavaScript يحتوي على معلومات حول حدث إجراء قد وقع. يستقبل مخزن Redux هذه الكائنات الإجرائية، ثم يُحدّث حالته بناءً عليها. أحيانًا يحمل إجراء Redux بعض البيانات، مثلًا يحمل الإجراء اسم المستخدم بعد تسجيل الدخول. رغم أن البيانات اختيارية، يجب أن تحمل الإجراءات خاصية `type` التي تحدد نوع الإجراء الذي وقع.

فكر في إجراءات Redux كرسل تنقل معلومات عن الأحداث التي تحدث في تطبيقك إلى مخزن Redux. ثم يقوم المخزن بإدارة تحديث الحالة بناءً على الإجراء الذي وقع.

# --instructions--

كتابة إجراء Redux بسيطة مثل إعلان كائن يحتوي على خاصية type. أعلن كائنًا باسم `action` وأعطه خاصية `type` بقيمة السلسلة النصية `'LOGIN'`.

# --hints--

يجب أن يوجد كائن `action`.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

يجب أن يحتوي كائن `action` على خاصية مفتاحية `type` بقيمة `LOGIN`.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
