---
id: 5a24c314108439a4d4036158
title: لا تُغيّر الحالة أبدًا
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

تصف هذه التحديات النهائية عدة طرق لفرض المبدأ الأساسي لعدم تغيير الحالة في Redux. الحالة غير القابلة للتغيير تعني أنك لا تعدل الحالة مباشرة، بل تعيد نسخة جديدة من الحالة.

إذا التقطت لقطة لحالة تطبيق Redux مع مرور الوقت، سترى شيئًا مثل `state 1`، `state 2`، `state 3`، `state 4`، `...` وهكذا، حيث قد تكون كل حالة مشابهة للأخرى، لكنها قطعة بيانات مستقلة. هذه الخاصية، في الواقع، هي التي توفر ميزات مثل تصحيح الأخطاء عبر الزمن التي قد تكون سمعت عنها.

لا يفرض Redux بشكل نشط عدم تغيير الحالة في المخزن أو المخفضات، فهذه المسؤولية تقع على المبرمج. لحسن الحظ، توفر JavaScript (وخاصة ES6) عدة أدوات مفيدة يمكنك استخدامها لفرض عدم تغيير الحالة، سواء كانت `string` أو `number` أو `array` أو `object`. لاحظ أن السلاسل النصية والأرقام هي قيم بدائية وغير قابلة للتغيير بطبيعتها. بمعنى آخر، الرقم 3 هو دائمًا 3. لا يمكنك تغيير قيمة الرقم 3. أما `array` أو `object` فهي قابلة للتغيير. في الممارسة، ستتكون حالتك على الأرجح من `array` أو `object`، إذ تعد هذه هياكل بيانات مفيدة لتمثيل أنواع كثيرة من المعلومات.

# --instructions--

يوجد في محرر الشفرة `store` و `reducer` لإدارة عناصر المهام. أكمل كتابة حالة `ADD_TO_DO` في المخفض لإضافة مهمة جديدة إلى الحالة. هناك عدة طرق لتحقيق ذلك باستخدام JavaScript القياسية أو ES6. حاول أن تجد طريقة لإرجاع مصفوفة جديدة مع إضافة العنصر من `action.todo` في النهاية.

# --hints--

يجب أن يوجد مخزن Redux وأن يبدأ بحالة تساوي مصفوفة `todos` في محرر الشفرة.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

يجب أن تكون كل من `addToDo` و `immutableReducer` دالتين.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

يجب أن يؤدي إرسال إجراء من النوع `ADD_TO_DO` على مخزن Redux إلى إضافة عنصر `todo` ويجب ألا يغير الحالة.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
