---
id: 5a24c314108439a4d403614f
title: إرسال حدث إجراء
challengeType: 6
forumTopicId: 301442
dashedName: dispatch-an-action-event
---

# --description--

طريقة `dispatch` هي التي تستخدمها لإرسال الإجراءات إلى مخزن Redux. استدعاء `store.dispatch()` وتمرير القيمة التي تُرجع من منشئ الإجراء يرسل إجراءً إلى المخزن.

تذكّر أن منشئي الإجراءات يُرجعون كائنًا يحتوي على خاصية type تحدد نوع الإجراء الذي حدث. ثم تُرسل الطريقة كائن الإجراء إلى مخزن Redux. بناءً على مثال التحدي السابق، الأسطر التالية متكافئة، وكلاهما يرسل الإجراء من النوع `LOGIN`:

```js
store.dispatch(actionCreator());
store.dispatch({ type: 'LOGIN' });
```

# --instructions--

مخزن Redux في محرر الشفرة يحتوي على حالة مبدئية وهي كائن يحتوي على خاصية `login` مضبوطة حاليًا على `false`. هناك أيضًا منشئ إجراء يسمى `loginAction()` يُرجع إجراءً من نوع `LOGIN`. أرسل إجراء `LOGIN` إلى مخزن Redux باستدعاء طريقة `dispatch`، ومرر الإجراء الذي أنشأه `loginAction()`.

# --hints--

استدعاء الدالة `loginAction` يجب أن يُرجع كائنًا بخاصية `type` مضبوطة على السلسلة النصية `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

يجب تهيئة المخزن بكائن يحتوي على خاصية `login` مضبوطة على `false`.

```js
assert(store.getState().login === false);
```

يجب استخدام طريقة `store.dispatch()` لإرسال إجراء من نوع `LOGIN`.

```js
assert(
    (function () {
      let noWhiteSpace = code.replace(/\s/g, '');
      return (
        noWhiteSpace.includes('store.dispatch(loginAction())') ||
        noWhiteSpace.includes("store.dispatch({type: 'LOGIN'})") === true
      );
    })()
);
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

// Dispatch the action here:
```

# --solutions--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

store.dispatch(loginAction());
```
