---
id: 5a24c314108439a4d4036159
title: استخدام معامل النشر على المصفوفات
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

أحد الحلول في ES6 للمساعدة في فرض عدم تحوُّل الحالة في Redux هو معامل النشر: `...`. يمتلك معامل النشر تطبيقات متعددة، وأحدها مناسب تمامًا للتحدي السابق المتمثل في إنتاج مصفوفة جديدة من مصفوفة موجودة. هذه صيغة نحوية جديدة نسبيًا، لكنها مستخدمة بشكل شائع. على سبيل المثال، إذا كان لديك مصفوفة `myArray` وكتبت:

```js
let newArray = [...myArray];
```

الآن `newArray` نسخة مستنسخة من `myArray`. كلا المصفوفتين موجودتان بشكل منفصل في الذاكرة. إذا أجريت تحوُّلًا مثل `newArray.push(5)`، فإن `myArray` لا تتغير. المعامل `...` يقوم فعليًا بـ *نشر* القيم في `myArray` داخل مصفوفة جديدة. لاستنساخ مصفوفة مع إضافة قيم إضافية في المصفوفة الجديدة، يمكنك كتابة `[...myArray, 'new value']`. هذا يعيد مصفوفة جديدة مكونة من قيم `myArray` والسلسلة النصية `new value` كآخر قيمة. يمكن استخدام صيغة النشر عدة مرات في تكوين المصفوفة بهذه الطريقة، لكن من المهم ملاحظة أنها تنشئ نسخة سطحية فقط من المصفوفة. أي أنها توفر عمليات مصفوفة غير قابلة للتحوُّل فقط للمصفوفات أحادية البعد.

# --instructions--

استخدم معامل النشر لإرجاع نسخة جديدة من الحالة عند إضافة مهمة جديدة.

# --hints--

يجب أن يوجد مخزن Redux ويبدأ بحالة تساوي `["Do not mutate state!"]`.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

يجب أن تكون كل من `addToDo` و`immutableReducer` دالتين.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

يجب أن يؤدي إرسال إجراء من النوع `ADD_TO_DO` على مخزن Redux إلى إضافة عنصر `todo` ويجب ألا يحوِّل الحالة.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

يجب استخدام معامل النشر لإرجاع الحالة الجديدة.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
