---
id: 5a24c314108439a4d4036157
title: كتابة عداد باستخدام Redux
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

الآن تعلمت كل المبادئ الأساسية لـ Redux! رأيت كيف تنشئ الإجراءات ومنشئي الإجراءات، تنشئ مخزن Redux، ترسل إجراءاتك إلى المخزن، وتصمم تحديثات الحالة باستخدام المخفضات النقية. حتى أنك رأيت كيف تدير الحالة المعقدة بتكوين المخفضات وتعالج الإجراءات غير المتزامنة. هذه الأمثلة بسيطة، لكن هذه المفاهيم هي المبادئ الأساسية لـ Redux. إذا فهمتها جيدًا، فأنت جاهز لبدء بناء تطبيق Redux الخاص بك. التحديات التالية تغطي بعض التفاصيل المتعلقة بعدم قابلية تغيير `state`، لكن أولًا، إليك مراجعة لكل ما تعلمته حتى الآن.

# --instructions--

في هذا الدرس، ستنفذ عدادًا بسيطًا باستخدام Redux من الصفر. الأساسيات متوفرة في محرر الشفرة، لكن عليك ملء التفاصيل! استخدم الأسماء المقدمة وعرّف منشئي الإجراءات `incAction` و `decAction`، ودالة `counterReducer()`، وأنواع الإجراءات `INCREMENT` و `DECREMENT`، وأخيرًا مخزن Redux `store`. بعد الانتهاء، يجب أن تكون قادرًا على إرسال إجراءات `INCREMENT` أو `DECREMENT` لزيادة أو تقليل الحالة المحتفظ بها في `store`. حظًا موفقًا في بناء أول تطبيق Redux لك!

# --hints--

يجب أن يعيد منشئ الإجراء `incAction` كائن إجراء مع خاصية `type` تساوي قيمة `INCREMENT`

```js
assert(incAction().type === INCREMENT);
```

يجب أن يعيد منشئ الإجراء `decAction` كائن إجراء مع خاصية `type` تساوي قيمة `DECREMENT`

```js
assert(decAction().type === DECREMENT);
```

تشغيل `store.getState()` يجب أن يعيد رقمًا

```js
assert(typeof store.getState() === 'number');
```

يجب أن يبدأ مخزن Redux بحالة `state` قيمتها 0.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

إرسال `incAction` إلى مخزن Redux يجب أن يزيد `state` بمقدار 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

إرسال `decAction` إلى مخزن Redux يجب أن ينقص `state` بمقدار 1.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

يجب أن يكون `counterReducer` دالة

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
