---
id: 5a24c314108439a4d403614c
title: الحصول على الحالة من مخزن Redux
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

يوفر كائن مخزن Redux عدة طرق تتيح لك التفاعل معه. على سبيل المثال، يمكنك استرجاع الـ `state` الحالي المحتفظ به في كائن مخزن Redux باستخدام الطريقة `getState()`.

# --instructions--

أُعيد كتابة الشفرة من التحدي السابق بشكل أكثر إيجازًا في محرر الشفرة. استخدم `store.getState()` لاسترجاع الـ `state` من الـ `store`، وقم بتعيينها إلى متغير جديد اسمه `currentState`.

# --hints--

يجب أن يحتوي مخزن Redux على القيمة 5 كحالة أولية.

```js
assert(store.getState() === 5);
```

يجب أن يوجد متغير `currentState` ويُعيَّن له الحالة الحالية لمخزن Redux.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
