---
id: 5a24c314108439a4d4036153
title: تسجيل مستمع للمخزن
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

طريقة أخرى يمكنك الوصول إليها على كائن Redux `store` هي `store.subscribe()`. تتيح لك هذه الطريقة الاشتراك بدوال مستمعة للمخزن، والتي تُستدعى في كل مرة يُرسل فيها إجراء إلى المخزن. استخدام بسيط لهذه الطريقة هو الاشتراك بدالة في المخزن تقوم فقط بتسجيل رسالة في كل مرة يُستقبل فيها إجراء ويتم تحديث المخزن.

# --instructions--

اكتب دالة استدعاء راجع تزيد المتغير العام `count` في كل مرة يستقبل فيها المخزن إجراءً، ومرر هذه الدالة إلى طريقة `store.subscribe()`. سترى أن `store.dispatch()` تُستدعى ثلاث مرات متتالية، كل مرة تمرر مباشرة كائن إجراء. راقب مخرجات وحدة التحكم بين عمليات الإرسال لترى التحديثات تحدث.

# --hints--

يجب أن يؤدي إرسال إجراء `ADD` إلى المخزن إلى زيادة الحالة بمقدار `1`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

يجب أن يكون هناك دالة مستمعة مشتركة في المخزن باستخدام `store.subscribe`.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

يجب أن تستقبل `store.subscribe` دالة.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

يجب ألا تُستدعى الدالة التي تُمرر إلى `store.subscribe` مباشرة.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

يجب أن تزيد دالة الاستدعاء الراجعة لـ `store.subscribe` أيضًا المتغير العام `count` مع تحديث المخزن.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
