---
id: 5a24c314108439a4d4036156
title: استخدام البرمجيات الوسيطة لمعالجة الإجراءات غير المتزامنة
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

حتى الآن، تجنبت هذه التحديات مناقشة الإجراءات غير المتزامنة، لكنها جزء لا مفر منه من تطوير الويب. في مرحلة ما، ستحتاج إلى استدعاء نقاط نهاية غير متزامنة في تطبيق Redux الخاص بك، فكيف تتعامل مع هذه الأنواع من الطلبات؟ يوفر Redux برمجيات وسيطة مصممة خصيصًا لهذا الغرض، تُسمى برمجيات Redux Thunk الوسيطة. إليك وصفًا موجزًا لكيفية استخدامها مع Redux.

لإضافة برمجيات Redux Thunk الوسيطة، تمررها كمعلمة إلى `Redux.applyMiddleware()`. ثم تُمرر هذه العبارة كمعلمة اختيارية ثانية إلى دالة `createStore()`. ألقِ نظرة على الشفرة في أسفل المحرر لترى ذلك. بعد ذلك، لإنشاء إجراء غير متزامن، تُعيد دالة في منشئ الإجراء تأخذ `dispatch` كمعلمة. داخل هذه الدالة، يمكنك استدعاء الإجراءات وتنفيذ الطلبات غير المتزامنة.

في هذا المثال، يُحاكى طلب غير متزامن باستخدام استدعاء `setTimeout()`. من الشائع استدعاء إجراء قبل بدء أي سلوك غير متزامن حتى يعرف حالة التطبيق أن بعض البيانات تُطلب (يمكن أن تعرض هذه الحالة رمز تحميل، على سبيل المثال). ثم، بمجرد استلام البيانات، تستدعي إجراءً آخر يحمل البيانات كحمولة مع معلومات تفيد بأن الإجراء اكتمل.

تذكر أنك تمرر `dispatch` كمعلمة إلى منشئ الإجراء الخاص هذا. هذا ما ستستخدمه لاستدعاء إجراءاتك، فقط مرر الإجراء مباشرة إلى dispatch والبرمجيات الوسيطة تتولى الباقي.

# --instructions--

اكتب كلا الاستدعاءين في منشئ الإجراء `handleAsync()`. استدعي `requestingData()` قبل `setTimeout()` (الاستدعاء المحاكى لواجهة برمجة التطبيقات). ثم، بعد استلام البيانات (الافتراضية)، استدعي إجراء `receivedData()`، ومرر هذه البيانات إليه. الآن تعرف كيف تتعامل مع الإجراءات غير المتزامنة في Redux. كل شيء آخر يستمر في العمل كما كان من قبل.

# --hints--

يجب أن يُعيد منشئ الإجراء `requestingData` كائنًا من النوع الذي يساوي قيمة `REQUESTING_DATA`.

```js
assert(requestingData().type === REQUESTING_DATA);
```

يجب أن يُعيد منشئ الإجراء `receivedData` كائنًا من النوع الذي يساوي قيمة `RECEIVED_DATA`.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

يجب أن يكون `asyncDataReducer` دالة.

```js
assert(typeof asyncDataReducer === 'function');
```

يجب أن يؤدي استدعاء منشئ الإجراء `requestingData` إلى تحديث خاصية `state` في المخزن لتصبح fetching بقيمة `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

يجب أن يؤدي استدعاء `handleAsync` إلى استدعاء إجراء طلب البيانات ثم استدعاء إجراء استلام البيانات بعد تأخير.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
