---
id: 5a24c314108439a4d403614b
title: إنشاء مخزن Redux
challengeType: 6
forumTopicId: 301439
dashedName: create-a-redux-store
---

# --description--

Redux هو إطار عمل لإدارة الحالة يمكن استخدامه مع عدد من تقنيات الويب المختلفة، بما في ذلك React.

في Redux، يوجد كائن حالة واحد مسؤول عن الحالة الكاملة لتطبيقك. هذا يعني إذا كان لديك تطبيق React يحتوي على عشرة مكوِّنات، وكان لكل مكوِّن حالة محلية خاصة به، فإن الحالة الكاملة لتطبيقك ستُعرّف بواسطة كائن حالة واحد موجود في مخزن Redux `store`. هذه هي المبدأ المهم الأول الذي يجب فهمه عند تعلم Redux: مخزن Redux هو المصدر الوحيد للحقيقة فيما يتعلق بحالة التطبيق.

هذا يعني أيضًا أنه في أي وقت يريد أي جزء من تطبيقك تحديث الحالة، يجب أن يتم ذلك من خلال مخزن Redux. تدفق البيانات أحادي الاتجاه يجعل من السهل تتبع إدارة الحالة في تطبيقك.

# --instructions--

`store` في Redux هو كائن يحتفظ بـ `state` الخاص بالتطبيق ويديره. توجد طريقة باسم `createStore()` في كائن Redux، وتستخدمها لإنشاء الـ `store`. هذه الطريقة تأخذ دالة `reducer` كمعلمة مطلوبة. دالة `reducer` سيتم شرحها في تحدٍّ لاحق، وهي معرفة مسبقًا في محرر الكود. هذه الدالة ببساطة تستقبل `state` كمعلمة وتعيد `state`.

أعلن عن متغير `store` وقم بتعيينه إلى طريقة `createStore()`، مع تمرير `reducer` كمعلمة.

**ملاحظة:** الكود في المحرر يستخدم صيغة المعلمات الافتراضية في ES6 لتهيئة هذه الحالة بقيمة `5`. إذا لم تكن معتادًا على المعلمات الافتراضية، يمكنك الرجوع إلى <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/es6/set-default-parameters-for-your-functions" target="_blank" rel="noopener noreferrer nofollow">قسم ES6 في المنهج</a> الذي يغطي هذا الموضوع.

# --hints--

يجب أن يوجد مخزن Redux.

```js
assert(typeof store.getState === 'function');
```

يجب أن يحتوي مخزن Redux على قيمة 5 للحالة.

```js
assert(store.getState() === 5);
```

# --seed--

## --seed-contents--

```js
const reducer = (state = 5) => {
  return state;
};

// Redux methods are available from a Redux object
// For example: Redux.createStore()
// Define the store here:
```

# --solutions--

```js
const reducer = (state = 5) => {
  return state;
}

const store = Redux.createStore(reducer);
```
