---
id: 5a24c314108439a4d4036155
title: إرسال بيانات الإجراء إلى المخزن
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

حتى الآن تعلمت كيف ترسل الإجراءات إلى مخزن Redux، لكن هذه الإجراءات لم تحتوي على أي معلومات سوى `type`. يمكنك أيضًا إرسال بيانات محددة مع إجراءاتك. في الواقع، هذا شائع جدًا لأن الإجراءات عادةً ما تنشأ من تفاعل المستخدم وتميل إلى حمل بعض البيانات معها. غالبًا ما يحتاج مخزن Redux إلى معرفة هذه البيانات.

# --instructions--

هناك `notesReducer()` أساسي ومنشئ إجراء `addNoteText()` معرفان في محرر الشفرة. أكمل جسم دالة `addNoteText()` بحيث تُرجع كائن `action`. يجب أن يتضمن الكائن خاصية `type` بقيمة `ADD_NOTE`، وأيضًا خاصية `text` تُعيّن إلى بيانات `note` التي تُمرر إلى منشئ الإجراء. عند استدعاء منشئ الإجراء، ستمرر معلومات ملاحظة محددة يمكنك الوصول إليها لإنشاء الكائن.

بعد ذلك، أكمل كتابة عبارة `switch` داخل الدالة `notesReducer()`. عليك إضافة حالة تتعامل مع الإجراءات الناتجة عن `addNoteText()`. يجب أن تُنفذ هذه الحالة كلما كان هناك إجراء من النوع `ADD_NOTE`، ويجب أن تُعيد خاصية `text` من الكائن `action` القادم كقيمة جديدة لـ `state`.

يُرسل الإجراء في أسفل الشفرة. بعد الانتهاء، شغّل الشفرة وراقب وحدة التحكم. هذا كل ما يلزم لإرسال بيانات محددة للإجراء إلى المخزن واستخدامها عند تحديث `state` في المخزن.

# --hints--

يجب أن يُرجع منشئ الإجراء `addNoteText` كائنًا يحتوي على المفاتيح `type` و `text`.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

يجب أن يؤدي إرسال إجراء من نوع `ADD_NOTE` باستخدام منشئ الإجراء `addNoteText` إلى تحديث `state` إلى السلسلة النصية التي تم تمريرها إلى منشئ الإجراء.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
