---
id: 5a24c314108439a4d403614e
title: تعريف منشئ إجراء
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

بعد إنشاء إجراء، الخطوة التالية هي إرسال الإجراء إلى مخزن Redux ليُحدّث حالته. في Redux، تُعرّف منشئي الإجراءات لتحقيق ذلك. منشئ الإجراء هو ببساطة دالة JavaScript تُرجع إجراءً. بعبارة أخرى، منشئو الإجراءات يُنشئون كائنات تمثل أحداث الإجراءات.

# --instructions--

عرّف دالة باسم `actionCreator()` تُرجع كائن الـ `action` عند استدعائها.

# --hints--

يجب أن توجد دالة `actionCreator`.

```js
assert(typeof actionCreator === 'function');
```

تشغيل دالة `actionCreator` يجب أن يُرجع كائن الـ `action`.

```js
assert(typeof action === 'object');
```

يجب أن يحتوي الـ `action` المُرجع على خاصية رئيسية `type` بقيمة `LOGIN`.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
