---
id: 5a24c314108439a4d4036154
title: دمج عدة مخفضات
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

عندما يبدأ حالة تطبيقك في التعقيد، قد تميل إلى تقسيم الحالة إلى أجزاء متعددة. بدلًا من ذلك، تذكر المبدأ الأول في Redux: كل حالة التطبيق تُحفظ في كائن حالة واحد داخل المخزن. لذلك، يوفر Redux تركيب المخفضات كحل لنموذج حالة معقد. تُعرّف عدة مخفضات للتعامل مع أجزاء مختلفة من حالة التطبيق، ثم تدمج هذه المخفضات معًا في مخفض جذر واحد. يُمرر المخفض الجذر بعد ذلك إلى دالة `createStore()` في Redux.

لكي نتمكن من دمج عدة مخفضات معًا، يوفر Redux طريقة `combineReducers()`. تقبل هذه الطريقة كائنًا كمعلمة تُعرّف فيه خصائص تربط المفاتيح بدوال مخفض محددة. الاسم الذي تعطيه للمفاتيح يستخدمه Redux كاسم للجزء المرتبط من الحالة.

عادةً، من الممارسات الجيدة إنشاء مخفض لكل جزء من حالة التطبيق عندما تكون مميزة أو فريدة بطريقة ما. على سبيل المثال، في تطبيق لتدوين الملاحظات مع توثيق المستخدم، يمكن لمخفض واحد التعامل مع التوثيق بينما يتعامل آخر مع النص والملاحظات التي يرسلها المستخدم. لتطبيق كهذا، قد نكتب طريقة `combineReducers()` بهذا الشكل:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

الآن، المفتاح `notes` سيحتوي على كل الحالة المرتبطة بملاحظاتنا والتي يتعامل معها `notesReducer`. هكذا يمكن تركيب عدة مخفضات لإدارة حالة تطبيق أكثر تعقيدًا. في هذا المثال، الحالة المحفوظة في مخزن Redux ستكون كائنًا واحدًا يحتوي على خاصيتي `auth` و`notes`.

# --instructions--

يوجد دالتي `counterReducer()` و`authReducer()` في محرر الشفرة، مع مخزن Redux. أكمل كتابة دالة `rootReducer()` باستخدام طريقة `Redux.combineReducers()`. عيّن `counterReducer` إلى مفتاح اسمه `count` و`authReducer` إلى مفتاح اسمه `auth`.

# --hints--

يجب أن يزيد ويقلل `counterReducer` من قيمة `state`.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

يجب أن يقوم `authReducer` بتبديل `state` الخاص بـ `authenticated` بين `true` و `false`.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

يجب أن يحتوي `state` الخاص بالمخزن على مفتاحين: `count` الذي يحمل رقمًا، و`auth` الذي يحمل كائنًا. يجب أن يحتوي كائن `auth` على خاصية `authenticated` التي تحمل قيمة منطقية.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

يجب أن تكون دالة `rootReducer` دالة تدمج بين `counterReducer` و`authReducer`.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
