---
id: 5a24c314108439a4d4036150
title: التعامل مع إجراء في المخزن
challengeType: 6
forumTopicId: 301444
dashedName: handle-an-action-in-the-store
---

# --description--

بعد إنشاء الإجراء وإرساله، يحتاج مخزن Redux لمعرفة كيفية الاستجابة لهذا الإجراء. هذه مهمة دالة `reducer`. في Redux، يكون الـ reducers مسؤولين عن تعديلات الحالة التي تحدث استجابةً للإجراءات. دالة `reducer` تستقبل `state` و `action` كمعلمتين، ودائمًا تعيد نسخة جديدة من `state`. من المهم أن تدرك أن هذا هو **الدور الوحيد** للـ reducer. لا توجد له أي آثار جانبية — لا يستدعي أي نقطة نهاية لواجهة برمجة التطبيقات ولا يحتوي على مفاجآت مخفية. الـ reducer هو ببساطة دالة نقية تستقبل state و action ثم تعيد state جديد.

مبدأ أساسي آخر في Redux هو أن `state` للقراءة فقط. بعبارة أخرى، يجب أن تُعيد الدالة `reducer` **دائمًا** نسخة جديدة من `state` وألّا تعدّل الحالة مباشرةً أبدًا. لا يفرض Redux ثبات الحالة، لكنك مسؤول عن فرضه في شفرة دوال reducer لديك. وستتدرب على هذا في تحديات لاحقة.

# --instructions--

يحتوي محرر الشفرة على المثال السابق وكذلك بداية دالة `reducer` لك. املأ جسم الدالة `reducer` بحيث إذا تلقّت إجراءً من النوع `'LOGIN'` تُعيد كائن حالة مع ضبط `login` على `true`. وإلا، فإنها تُعيد `state` الحالية. لاحظ أن `state` الحالية والإجراء `action` المُرسَل يُمرَّران إلى المخفِّض، بحيث يمكنك الوصول إلى نوع الإجراء مباشرةً عبر `action.type`.

# --hints--

يجب أن تُرجع استدعاء الدالة `loginAction` كائنًا بخاصية type مضبوطة على السلسلة النصية `LOGIN`.

```js
assert(loginAction().type === 'LOGIN');
```

يجب تهيئة المخزن بكائن يحتوي على الخاصية `login` مضبوطة على `false`.

```js
assert(store.getState().login === false);
```

يجب أن يؤدي إرسال `loginAction` إلى تحديث خاصية `login` في حالة المخزن إلى `true`.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginAction());
    const afterState = store.getState();
    return initialState.login === false && afterState.login === true;
  })()
);
```

إذا لم يكن الإجراء من النوع `LOGIN`، يجب أن يُرجع المخزن الحالة الحالية.

```js
assert(
  (function () {
    store.dispatch({ type: '__TEST__ACTION__' });
    let afterTest = store.getState();
    return typeof afterTest === 'object' && afterTest.hasOwnProperty('login');
  })()
);
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```

# --solutions--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {

  if (action.type === 'LOGIN') {
    return {login: true}
  }

  else {
    return state
  }

};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```
