---
id: 5a24c314108439a4d403615b
title: نسخ كائن باستخدام Object.assign
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

اشتغلت التحديات الأخيرة مع المصفوفات، لكن هناك طرق تساعد على فرض عدم تغير الحالة عندما تكون الحالة `object` أيضًا. أداة مفيدة للتعامل مع الكائنات هي الدالة `Object.assign()`. تأخذ `Object.assign()` كائن الهدف وكائنات المصدر وتُعيّن الخصائص من كائنات المصدر إلى كائن الهدف. تُستبدل أي خصائص متطابقة بالخصائص الموجودة في كائنات المصدر. يُستخدم هذا السلوك عادةً لعمل نسخ سطحية من الكائنات بتمرير كائن فارغ كالمعلمة الأولى يليه الكائن أو الكائنات التي تريد نسخها. إليك مثالًا:

```js
const newObject = Object.assign({}, obj1, obj2);
```

ينشئ هذا `newObject` كـ `object` جديد يحتوي على الخصائص الموجودة حاليًا في `obj1` و `obj2`.

# --instructions--

تم تعديل حالة Redux والإجراءات للتعامل مع `object` كـ `state`. حرر الشفرة لإرجاع كائن `state` جديد للإجراءات من النوع `ONLINE`، والذي يعيّن خاصية `status` إلى السلسلة النصية `online`. حاول استخدام `Object.assign()` لإكمال التحدي.

# --hints--

يجب أن يوجد مخزن Redux ويُهيأ بحالة تعادل كائن `defaultState` المُعلن في السطر 1.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp` و`immutableReducer` يجب أن يكونا دالتين.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

يجب أن يؤدي إرسال إجراء من النوع `ONLINE` إلى تحديث الخاصية `status` في الحالة إلى `online` ويجب ألا يُغير الحالة الأصلية.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

يجب استخدام `Object.assign` لإرجاع حالة جديدة.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
