---
id: 684aaf9ec670c68d20efd0d0
title: ユーザー設定マネージャーを作成する
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

このラボでは、テーマ、言語、通知などの設定を管理できるユーザー設定マネージャーを作成します。設定の追加、更新、削除、表示を行う関数を実装してください。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストを通過させてラボを完了させてください。

**ユーザーストーリー:**

1. 設定の辞書とキーと値のペアを含むタプルの2つのパラメータを持つ`add_setting`という名前の関数を定義してください。

1. `add_setting`関数は以下を行う必要があります：
    
   - キーと値を小文字に変換する。
   - キーの設定が存在する場合、`Setting '[key]' already exists! Cannot add a new setting with this name.`を返す。
   - キーの設定が存在しない場合、指定された設定の辞書にキーと値のペアを追加し、`Setting '[key]' added with value '[value]' successfully!`を返す。
   - 返されるメッセージはキーと値が小文字である必要があります。

1. 設定の辞書とキーと値のペアを含むタプルの2つのパラメータを持つ`update_setting`という名前の関数を定義してください。

1. `update_setting`関数は以下を行う必要があります：
       
   - キーと値を小文字に変換する。
   - キーの設定が存在する場合、指定された設定の辞書内で値を更新し、`Setting '[key]' updated to '[value]' successfully!`を返す。
   - キーの設定が存在しない場合、`Setting '[key]' does not exist! Cannot update a non-existing setting.`を返す。
   - 返されるメッセージはキーと値が小文字である必要があります。

1. 設定の辞書とキーの2つのパラメータを持つ`delete_setting`という名前の関数を定義してください。

1. `delete_setting`関数は以下を行う必要があります：
    
   - 渡されたキーを小文字に変換する。
   - キーの設定が存在する場合、指定された設定の辞書からキーと値のペアを削除し、`Setting '[key]' deleted successfully!`を返す。
   - キーの設定が存在しない場合、`Setting not found!`を返す。
   - 返されるメッセージはキーが小文字である必要があります。

1. 設定の辞書を1つのパラメータに持つ`view_settings`という名前の関数を定義してください。

1. `view_settings`関数は以下を行う必要があります：
   
   - 指定された設定の辞書が空の場合、`No settings available.`を返す。
   - 辞書に設定が含まれている場合、設定を表示する文字列を返す。文字列は`Current User Settings:`で始まり、キーと値のペアをそれぞれ新しい行に表示し、キーは大文字で始まる必要があります。例えば、`view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})`は以下を返すべきです：

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. コードのテスト用に、いくつかのユーザー設定の好みを保存する`test_settings`という名前の辞書を作成してください。

# --hints--

`test_settings`という名前の辞書を作成し、いくつかの値を追加してください。

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

`add_setting`という名前の関数を定義してください。

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

`add_setting`関数は2つのパラメータを持つ必要があります。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting`はキーを小文字に変換する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting`は値を小文字に変換する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))`はエラーメッセージ`Setting 'theme' already exists! Cannot add a new setting with this name.`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))`は新しいキーと値のペアを追加し、成功メッセージ`Setting 'volume' added with value 'high' successfully!`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting`は指定されたキーと値のペアを辞書に正しく追加する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

`update_setting`という名前の関数を定義してください。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

`update_setting`関数は2つのパラメータを持つ必要があります。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting`関数はキーを小文字に変換する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting`関数は値を小文字に変換する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))`は既存のキーを更新し、成功メッセージ`Setting 'theme' updated to 'dark' successfully!`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))`はキーが存在しない場合、エラーメッセージ`Setting 'volume' does not exist! Cannot update a non-existing setting.`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting`は指定されたキーと値のペアを辞書に正しく更新する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

`delete_setting`という名前の関数を定義してください。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

`delete_setting`関数は2つのパラメータを持つ必要があります。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting`はキーを小文字に変換する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')`は既存のキーを削除し、成功メッセージ`Setting 'theme' deleted successfully!`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting`はキーが存在しない場合、エラーメッセージ`Setting not found!`を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting`は指定されたキーを辞書から正しく削除する必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

`view_settings`という名前の関数を定義してください。

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

`view_settings`関数は1つのパラメータを持つ必要があります。

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings`は指定された辞書が空の場合、メッセージ`No settings available.`を返す必要があります。

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings`は空でない辞書に対してフォーマットされた設定を返す必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings`は各設定名の最初の文字を大文字にする必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings`は正しい結果を表示し、改行文字で終わる必要があります。

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
