---
id: 65c4ef8463f869d0eaf87c67
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在 `TreeNode` 类中，将 `pass` 替换为 `__init__` 方法，以便可以初始化对象的属性。 暂时不要添加任何参数。

# --hints--

你应该从 `TreeNode` 类中删除 `pass` 关键字，并将其移至 `__init__` 方法内。

```js
({
  test: () => {
    const pyClassStr = runPython(`str(_Node(_code).find_class("TreeNode"))`);
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});
```

你应该在 `TreeNode` 类中定义一个 `__init__` 方法。 请记住在其内部使用 `pass` 关键字。

```js
({
  test: () => {
    assert(
      runPython(`_Node(_code).find_class("TreeNode").has_function("__init__")`)
    );
  },
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class TreeNode:
    pass
--fcc-editable-region--
```
