---
id: 65c4f0b82db314d3fcc0b8e1
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

在 `BinarySearchTree` 类中，将 `pass` 替换为 `__init__` 方法，并向该方法添加一个 `self` 参数。

# --hints--

你应该从 `BinarySearchTree` 类中删除 `pass` 关键字。

```js
({
  test: () => {
    const pyClassStr = runPython(
      `str(_Node(_code).find_class("BinarySearchTree"))`
    );
    const to_test = pyClassStr.split("\n");
    assert.notInclude(to_test[1], "pass");
  },
});

```


你应该定义一个 `__init__` 方法并向该方法添加一个 `self` 参数。 请记住在方法体内使用 `pass` 关键字。

```js
({ test: () => assert.match(code, /^\s{4}def\s+__init__\s*\(\s*self\s*\)\s*:/m) })
```

# --seed--

## --seed-contents--

```py
class TreeNode:
    def __init__(self, key):
        self.key = key
        self.left = None
        self.right = None

--fcc-editable-region--
class BinarySearchTree:
    pass
--fcc-editable-region--
```
