---
id: 684aaf9ec670c68d20efd0d0
title: 사용자 설정 관리자 만들기
challengeType: 27
dashedName: build-a-user-configuration-manager
saveSubmissionToDB: true
---

# --description--

이 실습에서는 테마, 언어, 알림 같은 설정을 관리할 수 있는 사용자 설정 관리자를 만듭니다. 설정을 추가, 업데이트, 삭제, 조회하는 함수를 구현하세요.

**목표:** 아래 사용자 스토리를 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 설정 딕셔너리와 키-값 쌍을 담은 튜플 두 개를 매개변수로 받는 `add_setting`라는 함수를 정의하세요.

1. `add_setting` 함수는 다음을 수행해야 합니다:
    
   - 키와 값을 소문자로 변환하세요.
   - 키 설정이 존재하면 `Setting '[key]' already exists! Cannot add a new setting with this name.`를 반환하세요.
   - 키 설정이 없으면 주어진 설정 딕셔너리에 키-값 쌍을 추가하고 `Setting '[key]' added with value '[value]' successfully!`를 반환하세요.
   - 반환하는 메시지에는 키와 값이 소문자로 포함되어야 합니다.

1. 설정 딕셔너리와 키-값 쌍을 담은 튜플 두 개를 매개변수로 받는 `update_setting`라는 함수를 정의하세요.

1. `update_setting` 함수는 다음을 수행해야 합니다:
       
   - 키와 값을 소문자로 변환하세요.
   - 키 설정이 존재하면 주어진 설정 딕셔너리에서 값을 업데이트하고 `Setting '[key]' updated to '[value]' successfully!`를 반환하세요.
   - 키 설정이 없으면 `Setting '[key]' does not exist! Cannot update a non-existing setting.`를 반환하세요.
   - 반환하는 메시지에는 키와 값이 소문자로 포함되어야 합니다.

1. 설정 딕셔너리와 키 하나를 매개변수로 받는 `delete_setting`라는 함수를 정의하세요.

1. `delete_setting` 함수는 다음을 수행해야 합니다:
    
   - 전달된 키를 소문자로 변환하세요.
   - 키 설정이 존재하면 주어진 설정 딕셔너리에서 키-값 쌍을 삭제하고 `Setting '[key]' deleted successfully!`를 반환하세요.
   - 키 설정이 없으면 `Setting not found!`를 반환하세요.
   - 반환하는 메시지에는 키가 소문자로 포함되어야 합니다.

1. 설정 딕셔너리 하나를 매개변수로 받는 `view_settings`라는 함수를 정의하세요.

1. `view_settings` 함수는 다음을 수행해야 합니다:
   
   - 주어진 설정 딕셔너리가 비어 있으면 `No settings available.`를 반환하세요.
   - 딕셔너리에 설정이 있으면 설정을 표시하는 문자열을 반환하세요. 문자열은 `Current User Settings:`로 시작하고, 각 키-값 쌍은 새 줄에 키가 대문자로 시작하도록 표시하세요. 예를 들어 `view_settings({'theme': 'dark', 'notifications': 'enabled', 'volume': 'high'})`는 다음과 같이 반환해야 합니다:

   ```md
   Current User Settings:
   Theme: dark
   Notifications: enabled
   Volume: high
   
   ```

1. 코드를 테스트하기 위해 `test_settings`라는 딕셔너리를 만들어 사용자 설정 일부를 저장하세요.

# --hints--

`test_settings`라는 딕셔너리를 만들고 몇 가지 값을 추가하세요.

```js
({
    test: () => runPython(`
        is_dict = isinstance(test_settings, dict)
        length = len(test_settings)
        assert is_dict and length > 0
    `)
})
```

`add_setting`라는 함수를 정의하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("add_setting")`)) })
```

`add_setting` 함수는 두 개의 매개변수를 가져야 합니다.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(add_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`add_setting`는 키를 소문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("VOLUME", "high"))
        assert "volume" in settings_in_test
        assert "VOLUME" not in settings_in_test
    `)
})
```

`add_setting`는 값을 소문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "HIGH"))
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting({'theme': 'light'}, ('THEME', 'dark'))`는 오류 메시지 `Setting 'theme' already exists! Cannot add a new setting with this name.`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = add_setting(settings_in_test, ("THEME", "dark"))
        expected_msg = "Setting 'theme' already exists! Cannot add a new setting with this name."
        assert result == expected_msg
    `)
})
```

`add_setting({'theme': 'light'}, ('volume', 'high'))`는 새 키-값 쌍을 추가하고 성공 메시지 `Setting 'volume' added with value 'high' successfully!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' added with value 'high' successfully!"
        assert result == expected_msg
        assert settings_in_test["volume"] == "high"
    `)
})
```

`add_setting`는 주어진 키-값 쌍을 딕셔너리에 올바르게 추가해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = add_setting(settings_in_test, ("os", "Ubuntu"))
        assert "os" in settings_in_test
        assert settings_in_test["os"] == "ubuntu"
    `)
})
```

`update_setting`라는 함수를 정의하세요.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('update_setting')
    `)
})
```

`update_setting` 함수는 두 개의 매개변수를 가져야 합니다.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(update_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`update_setting` 함수는 키를 소문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("THEME", "dark"))
        assert "THEME" not in settings_in_test
        assert "theme" in settings_in_test
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting` 함수는 값을 소문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        # Use the existing "theme" key but an uppercase value
        result = update_setting(settings_in_test, ("theme", "HIGH"))
        assert settings_in_test["theme"] == "high"
    `)
})
```

`update_setting({'theme': 'light'}, ('theme', 'dark'))`는 기존 키를 업데이트하고 성공 메시지 `Setting 'theme' updated to 'dark' successfully!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("theme", "dark"))
        expected_msg = "Setting 'theme' updated to 'dark' successfully!"
        assert result == expected_msg
        assert settings_in_test["theme"] == "dark"
    `)
})
```

`update_setting({'theme': 'light'}, ('volume', 'high'))`는 키가 없을 때 오류 메시지 `Setting 'volume' does not exist! Cannot update a non-existing setting.`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = update_setting(settings_in_test, ("volume", "high"))
        expected_msg = "Setting 'volume' does not exist! Cannot update a non-existing setting."
        assert result == expected_msg
    `)
})
```

`update_setting`는 주어진 키-값 쌍을 딕셔너리에 올바르게 업데이트해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = update_setting(settings_in_test, ("language", "spanish"))
        assert "language" in settings_in_test
        assert settings_in_test["language"] == "spanish"
    `)
})
```

`delete_setting`라는 함수를 정의하세요.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('delete_setting')
    `)
})
```

`delete_setting` 함수는 두 개의 매개변수를 가져야 합니다.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(delete_setting)
        assert len(sig.parameters) == 2
    `)
})
```

`delete_setting`는 키를 소문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "THEME")
        assert "theme" not in settings_in_test
        assert "THEME" not in settings_in_test
    `)
})
```

`delete_setting({'theme': 'light'}, 'theme')`는 기존 키를 삭제하고 성공 메시지 `Setting 'theme' deleted successfully!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "theme")
        expected_msg = "Setting 'theme' deleted successfully!"
        assert result == expected_msg
        assert "theme" not in settings_in_test
    `)
})
```

`delete_setting`는 키가 없을 때 오류 메시지 `Setting not found!`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light"}
        result = delete_setting(settings_in_test, "volume")
        expected_msg = "Setting not found!"
        assert result == expected_msg
    `)
})
```

`delete_setting`는 주어진 키를 딕셔너리에서 올바르게 삭제해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "English"}
        result = delete_setting(settings_in_test, "language")
        assert "language" not in settings_in_test
    `)
})
```

`view_settings`라는 함수를 정의하세요.

```js
({
    test: () => runPython(`
        assert _Node(_code).has_function('view_settings')
    `)
})
```

`view_settings` 함수는 하나의 매개변수를 가져야 합니다.

```js
({
    test: () => runPython(`
        import inspect
        sig = inspect.signature(view_settings)
        assert len(sig.parameters) == 1
    `)
})
```

`view_settings`는 주어진 딕셔너리가 비어 있으면 메시지 `No settings available.`를 반환해야 합니다.

```js
({
    test: () => runPython(`
        empty_settings = {}
        result = view_settings(empty_settings)
        expected_msg = "No settings available."
        assert result == expected_msg
    `)
})
```

`view_settings`는 비어 있지 않은 딕셔너리에 대해 포맷된 설정을 반환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "notifications": "enabled", "volume": "high"}
        result = view_settings(settings_in_test)
        assert result.startswith("Current User Settings:")
        assert "Theme: dark" in result
        assert "Notifications: enabled" in result
        assert "Volume: high" in result
    `)
})
```

`view_settings`는 각 설정 이름의 첫 글자를 대문자로 변환해야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "dark", "volume": "high"}
        result = view_settings(settings_in_test)
        assert "Theme:" in result
        assert "Volume:" in result
        assert "theme:" not in result
        assert "volume:" not in result
    `)
})
```

`view_settings`는 올바른 결과를 표시하고 줄 바꿈 문자로 끝나야 합니다.

```js
({
    test: () => runPython(`
        settings_in_test = {"theme": "light", "language": "english", "notifications": "enabled"}
        result = view_settings(settings_in_test)
        assert result == "Current User Settings:\\nTheme: light\\nLanguage: english\\nNotifications: enabled\\n"
    `)
})

```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def add_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        return f"Setting '{key}' already exists! Cannot add a new setting with this name."
    else:
        settings[key] = value
        return f"Setting '{key}' added with value '{value}' successfully!"

def update_setting(settings, key_value):
    key, value = key_value
    key = key.lower()
    value = value.lower()
    if key in settings:
        settings[key] = value
        return f"Setting '{key}' updated to '{value}' successfully!"
    else:
        return f"Setting '{key}' does not exist! Cannot update a non-existing setting."

def delete_setting(settings, key):
    key = key.lower()
    if key in settings:
        del settings[key]
        return f"Setting '{key}' deleted successfully!"
    else:
        return "Setting not found!"

def view_settings(settings):
    if not settings:
        return "No settings available."
    else:
        settings_str = "Current User Settings:\n"
        for key, value in settings.items():
            settings_str += f"{key.capitalize()}: {value}\n"
        return settings_str

test_settings = {
    "theme": "light",
    "language": "english",
    "notifications": "enabled"
}

print(add_setting(test_settings, ("VOLUME", "HIGH")))
print(update_setting(test_settings, ("theme", "dark")))
print(delete_setting(test_settings, "language"))
print(view_settings(test_settings))
```
