---
id: 66adc6046acc18453decf577
title: 12단계
challengeType: 1
dashedName: step-12
---

# --description--

봇의 새 별명을 보려면 콘솔에 새 메시지를 기록해야 합니다.

`newNicknameGreeting`라는 새 변수를 만드세요.

그런 다음 문자열 연결을 사용해 `"I love my nickname but I wish people would call me "` 문자열과 `bot` 변수를 마침표와 함께 이어 붙이세요.

결과를 `newNicknameGreeting` 변수에 할당하세요.

그다음 `newNicknameGreeting`의 값을 콘솔에 기록하세요.

# --hints--

`newNicknameGreeting`라는 변수가 있어야 합니다.

```js
assert.isNotNull(newNicknameGreeting);
```

`newNicknameGreeting` 변수는 문자열이어야 합니다.

```js
assert.isString(newNicknameGreeting);
```

문자열 연결을 사용해 `"I love my nickname but I wish people would call me "` 문자열과 `bot` 변수를 마침표와 함께 이어 붙여야 합니다.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

코드에 `console` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

`newNicknameGreeting`의 값을 콘솔에 기록해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
