---
id: 66ad8ab945d266383f318cbf
title: 2단계
challengeType: 1
dashedName: step-2
---

# --description--

이제 콘솔에서 봇의 첫 번째 메시지를 볼 수 있을 것입니다.

봇의 두 번째 메시지를 추가할 시간입니다.

콘솔에 메시지 `console.log`를 출력하는 또 다른 `"I am excited to talk to you."` 문장을 추가하세요.

# --hints--

코드에 두 번째 `console.log()` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

`console` 문장에는 메시지 `"I am excited to talk to you."`가 있어야 합니다. 따옴표를 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
