---
id: 66adcf383276814776aba3ca
title: 14단계
challengeType: 1
dashedName: step-14
---

# --description--

다음으로, `favoriteSubjectSentence`라는 변수를 만드세요.

문자열 연결을 사용해 문자열 `"My favorite subject is "`과 `favoriteSubject` 변수를 마침표와 함께 이어 붙이세요.

그 결과를 `favoriteSubjectSentence` 변수에 할당하세요.

그런 다음, `favoriteSubjectSentence`의 값을 콘솔에 출력하세요.

# --hints--

`favoriteSubjectSentence`라는 변수가 있어야 합니다.

```js
assert.isNotNull(favoriteSubjectSentence);
```

`favoriteSubjectSentence` 변수는 문자열이어야 합니다.

```js
assert.isString(favoriteSubjectSentence);
```

문자열 `"My favorite subject is "`과 `favoriteSubject` 변수를 마침표와 함께 문자열 연결로 이어 붙여야 합니다. 공백에 주의하세요.

```js
assert.strictEqual(favoriteSubjectSentence, "My favorite subject is Computer Science.");
```

코드에 `console` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 8);
```

`favoriteSubjectSentence`의 값을 콘솔에 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*favoriteSubjectSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

--fcc-editable-region--

--fcc-editable-region--
```
