---
id: 66ad8d150264db3926eccfeb
title: 3단계
challengeType: 1
dashedName: step-3
---

# --description--

이전 수업에서 `let` 키워드와 JavaScript에서 변수를 선언하는 방법을 배웠습니다.

다음은 `let` 키워드를 사용해 변수를 선언하는 방법을 다시 알려드립니다:

```js
let greeting;
```

`let` 키워드를 사용해 `bot`라는 변수를 선언하세요.

**참고**: 여기서는 `let`를 사용합니다. 워크숍 후반부에 `bot` 변수의 값을 변경할 예정이기 때문입니다.

# --hints--

코드에 `let` 키워드가 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let/g);
```

`bot`라는 변수를 선언해야 합니다.

```js
assert.isNotNull(bot);
```

변수를 선언할 때 `let` 키워드를 사용해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /let\s+bot;?/g);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
--fcc-editable-region--

--fcc-editable-region--
```
