---
id: 66add47d27763c4862492c8c
title: 15단계
challengeType: 1
dashedName: step-15
---

# --description--

마지막 단계로, `"Well, it was nice to talk to you. Have a nice day!"`의 봇 메시지를 콘솔에 기록하세요.

이로써 인사 봇이 완성됩니다!

# --before-each--

```js
const logSpy = __helpers.spyOn(console, 'log');
```

# --hints--

코드에 `console` 문장이 있어야 합니다.

```js
assert.lengthOf(logSpy.calls, 9);
```

`console` 문장은 문자열 `"Well, it was nice to talk to you. Have a nice day!"`을 기록해야 합니다.

```js
assert.deepInclude(logSpy.calls, ["Well, it was nice to talk to you. Have a nice day!"]);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

const favoriteSubjectSentence = "My favorite subject is " + favoriteSubject + ".";
console.log(favoriteSubjectSentence);

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

const favoriteSubject = "Computer Science";

const favoriteSubjectSentence = "My favorite subject is " + favoriteSubject + ".";
console.log(favoriteSubjectSentence);

console.log("Well, it was nice to talk to you. Have a nice day!");
```
