---
id: 66adb844118ba74107ce771f
title: 8단계
challengeType: 1
dashedName: step-8
---

# --description--

다음 봇 메시지는 봇의 위치에 관한 내용입니다.

`botLocationSentence`라는 변수를 만드세요.

그다음 `+` 연산자를 사용해 문자열 연결을 하여 `"I live in "` 문자열과 `botLocation` 변수를 마침표(`.`)와 함께 이어 붙이세요.

이 값을 `botLocationSentence` 변수에 할당하세요.

그다음 `botLocationSentence` 변수의 값을 콘솔에 출력하세요.

# --hints--

`botLocationSentence` 변수가 있어야 합니다.

```js
assert.isNotNull(botLocationSentence);
```

`botLocationSentence` 변수는 문자열 값을 가져야 합니다.

```js
assert.isString(botLocationSentence);
```

`+` 연산자를 사용해 문자열 `"I live in "`과 `botLocation` 변수를 마침표(`.`)와 함께 연결해야 합니다. 공백에 주의하세요.

```js
assert.equal(botLocationSentence, "I live in the universe.");
```

코드에 `console` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 5);
```

`console` 문장은 `botLocationSentence` 변수를 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botLocationSentence\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
