---
id: 66adc82170acb1464bd348eb
title: 13단계
challengeType: 1
dashedName: step-13
---

# --description--

봇의 마지막 몇 메시지는 봇이 가장 좋아하는 주제에 포커스를 맞출 것입니다.

`favoriteSubject`라는 새 변수를 만들고 문자열 `"Computer Science"`을 할당하는 것부터 시작하세요.

# --hints--

`favoriteSubject`라는 변수가 있어야 합니다.

```js
assert.isNotNull(favoriteSubject);
```

`favoriteSubject` 변수는 문자열이어야 합니다.

```js
assert.isString(favoriteSubject);
```

문자열 `"Computer Science"`를 `favoriteSubject` 변수에 할당해야 합니다.

```js
assert.strictEqual(favoriteSubject, "Computer Science");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

const newNicknameGreeting = "I love my nickname but I wish people would call me " + bot + ".";
console.log(newNicknameGreeting);

--fcc-editable-region--

--fcc-editable-region--
```
