---
id: 66ad9c633fa26d3c1475eae3
title: 6단계
challengeType: 1
dashedName: step-6
---

# --description--

이제 봇의 또 다른 메시지를 추가할 차례입니다.

메시지 `console`를 기록하는 또 다른 `"Allow me to introduce myself."` 문장을 코드에 추가하세요.

# --hints--

코드에 세 번째 `console.log()` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

`console` 문장에는 메시지 `"Allow me to introduce myself."`가 있어야 합니다. 따옴표를 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
