---
id: 66adc42868cab843ccee87d9
title: 10단계
challengeType: 1
dashedName: step-10
---

# --description--

이제 새로운 `bot` 값을 확인할 시간입니다.

먼저 `nicknameIntroduction`이라는 새 변수를 만드세요.

문자열 연결을 사용해 `"My nickname is "` 문자열과 `bot` 변수를 마침표(`.`)와 함께 이어 붙이세요.

그 결과 문자열을 `nicknameIntroduction` 변수에 할당하세요.

그다음 `nicknameIntroduction` 값을 콘솔에 출력하세요.

# --hints--

`nicknameIntroduction`라는 변수가 있어야 합니다.

```js
assert.isNotNull(nicknameIntroduction);
```

`nicknameIntroduction` 변수는 문자열이어야 합니다.

```js
assert.isString(nicknameIntroduction);
```

문자열 `"My nickname is "`과 `bot` 변수를 마침표(`.`)와 함께 문자열 연결로 이어 붙이세요. 공백에 주의하세요.

```js
assert.strictEqual(nicknameIntroduction, "My nickname is professorBot.");
```

코드에 `console` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 6);
```

`nicknameIntroduction` 값을 콘솔에 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*nicknameIntroduction\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

--fcc-editable-region--

--fcc-editable-region--
```
