---
id: 66ada3f46945763dd97f43f8
title: 7단계
challengeType: 1
dashedName: step-7
---

# --description--

이전에 `bot`와 `botLocation` 변수를 만들었습니다. 이제 이 변수들을 사용해 콘솔에 새로운 메시지를 출력할 것입니다.

이전 수업에서 `+` 연산자를 사용해 문자열을 연결하는 방법을 배웠습니다. 문자열을 이렇게 연결합니다:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

문자열과 변수를 연결할 때 공백에 주의해야 한다는 점을 기억하세요.

`botIntroduction`라는 변수를 만드세요.

그다음 `+` 연산자를 사용해 문자열 `"My name is "`와 `bot` 변수를 이어 붙이고, 마침표(`.`)를 덧붙이세요.

이 값을 `botIntroduction` 변수에 할당하세요.

마지막으로 `botIntroduction` 변수를 콘솔에 출력하세요.

# --hints--

`botIntroduction`라는 변수가 있어야 합니다.

```js
assert.isNotNull(botIntroduction);
```

`botIntroduction` 변수는 문자열 값을 가져야 합니다.

```js
assert.isString(botIntroduction);
```

문자열 `+`과 `"My name is "` 변수를 `bot` 연산자로 연결하고, 마침표(`.`)를 덧붙이세요. 공백에 주의하세요. 이 연결된 문자열을 `botIntroduction` 변수에 할당하세요.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

코드에 네 번째 `console.log()` 문장이 있어야 합니다.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

`console` 문장은 `botIntroduction` 변수를 출력해야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
