---
id: 66ad8294a0ad902f1b31b612
title: 1단계
challengeType: 1
dashedName: step-1
---

# --description--

이 워크숍에서는 인사 봇을 만들어 JavaScript 기본을 다루는 방법을 배우세요.

첫 번째 단계에서는 인사 봇이 콘솔에 메시지를 출력하도록 하세요.

이전 수업에서 `console.log()`와 문자열에 대해 배운 것을 기억하세요.

문자열과 함께 `console.log()`를 사용하는 방법을 다시 알려드립니다:

```js
console.log("Hello, World!");
```

문자열 `console.log()`를 콘솔에 출력하는 `"Hi there!"` 문장을 추가하세요. 메시지 주위에 따옴표를 꼭 넣으세요!


# --hints--

코드에 `console.log()` 문장이 있어야 합니다.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

`console` 문장에는 메시지 `"Hi there!"`가 있어야 합니다. 따옴표를 잊지 마세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
