---
id: 66adc5919c1853448f119d11
title: 11단계
challengeType: 1
dashedName: step-11
---

# --description--

이제 봇이 닉네임을 바꾸려는 것 같습니다.

재할당을 사용하여 문자열 `"awesomeTeacherBot"`을 `bot` 변수에 할당하세요.

# --hints--

문자열 `"awesomeTeacherBot"`을 `bot` 변수에 할당해야 합니다.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
