---
id: 66ad998285df023b23dacdd3
title: 5단계
challengeType: 1
dashedName: step-5
---

# --description--

이제 `botLocation` 변수를 초기화할 시간입니다.

여러 단어로 된 변수를 선언해야 할 때는 <dfn>카멜 케이스</dfn> 명명 관례를 사용할 수 있습니다.

`camelCase`를 사용할 때 첫 단어는 모두 소문자이고, 그 다음 단어들의 첫 글자는 대문자로 씁니다.

예시는 다음과 같습니다:

```js
let thisIsCamelCase = "Alice the camel is actually a horse.";
```

`"the universe"` 키워드를 사용해 같은 줄에서 `botLocation` 변수에 문자열 `let`를 선언하고 할당하세요.

# --hints--

변수를 선언하려면 `let` 키워드를 사용하세요.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/let/g), 2);
```

변수 이름은 `botLocation`이어야 합니다.

```js
assert.isNotNull(botLocation);
```

`botLocation` 변수에 문자열을 할당해야 합니다.

```js
assert.isString(botLocation);
```

`botLocation` 변수에 문자열 `"the universe"`가 할당되어야 합니다.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
