---
id: 66adbc8a9793d64250f7e609
title: 9단계
challengeType: 1
dashedName: step-9
---

# --description--

이전 수업에서 변수에 값을 다시 할당하는 방법을 배웠습니다:

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

다시 할당을 사용하여 문자열 `"professorBot"`을 `bot` 변수에 할당하세요.

# --hints--

문자열 `"professorBot"`을 `bot` 변수에 할당해야 합니다.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
bot = "teacherBot";

let botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
