---
id: 66ad984e92ce213a66f7e51c
title: 4단계
challengeType: 1
dashedName: step-4
---

# --description--

이제 `bot` 변수에 값을 할당할 차례입니다.

이전 수업에서 변수에 값을 할당하는 방법을 배웠습니다:

```js
variableName = "Here is the value";
```

대입 연산자 `=` 오른쪽에 있는 것이 왼쪽 변수에 할당하는 값임을 기억하세요.

문자열 `"teacherBot"` 를 `bot` 변수에 할당하세요.

# --hints--

`bot` 변수에 문자열을 할당해야 합니다.

```js
assert.isString(bot);
```

`bot` 변수에는 문자열 `"teacherBot"` 가 할당되어야 합니다.

```js
assert.equal(bot, "teacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");
let bot;
--fcc-editable-region--

--fcc-editable-region--
```
