---
id: 66ad998285df023b23dacdd3
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora, é hora de atribuir alguns valores às variáveis `bot` e `botLocation`.

In the previous lectures, you learned how to assign values to variables like this:

```js
variableName = "Here is the value";
```

Lembre-se: o que está do lado direito do operador de atribuição `=` é o valor que você está atribuindo à variável do lado esquerdo.

Atribua a string `"teacherBot"` à variável `bot` e a string `"the universe"` à variável `botLocation`.

# --hints--

Você deve atribuir uma string à variável `bot`.

```js
assert.isString(bot);
```

Você deve atribuir à variável `bot` a string `"teacherBot"`.

```js
assert.equal(bot, "teacherBot");
```

Você deve atribuir uma string à variável `botLocation`.

```js
assert.isString(botLocation);
```

Você deve atribuir à variável `botLocation` a string `"the universe"`.

```js
assert.equal(botLocation, "the universe");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

--fcc-editable-region--

--fcc-editable-region--
```
