---
id: 66ad9c633fa26d3c1475eae3
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

Agora, é hora de adicionar outra mensagem do bot.

Adicione outra instrução de `console` ao código que registra a mensagem `"Allow me to introduce myself."`.

# --hints--

Você deve ter uma instrução de `console`.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 3);
```

A instrução de `console` deve ter a mensagem `"Allow me to introduce myself."`. Não se esqueça das aspas.

```js
assert.match(__helpers.removeJSComments(code), /("|')Allow\s+me\s+to\s+introduce\s+myself.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

--fcc-editable-region--

--fcc-editable-region--
```
