---
id: 66ad8294a0ad902f1b31b612
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você aprenderá a trabalhar com os fundamentos de JavaScript criando um bot de saudações.

Neste primeiro passo, exiba uma mensagem do bot de saudações no console.

Remember that you learned about `console.log()` and strings in the previous lectures.

Aqui está um lembrete de como usar `console.log()` com strings:

```js
console.log("Hello, World!");
```

Adicione uma instrução de `console.log()` que exiba a string `"Hi there!"` no console. Não se esqueça das aspas ao redor da mensagem!


# --hints--

Você deve ter uma instrução de `console.log()` no código.

```js
assert.match(__helpers.removeJSComments(code), /console\.log(.*)/);
```

A instrução `console` deve ter a mensagem `"Hi there!"`. Não se esqueça das aspas.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\((['"])Hi\s+there!\1\);?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
