---
id: 66adbc8a9793d64250f7e609
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Nos vídeos das aulas anteriores, você aprendeu como reatribuir valores a variáveis assim:

```js
let name = "John";
console.log(name); // "John"

name = "Jane";
console.log(name); // "Jane"
```

Usando reatribuição, atribua a string `"professorBot"` à variável `bot`.

# --hints--

Você deve atribuir a string `"professorBot"` à variável `bot`.

```js
assert.strictEqual(bot, "professorBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

--fcc-editable-region--

--fcc-editable-region--
```
