---
id: 66ada3f46945763dd97f43f8
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Anteriormente, você criou as variáveis `bot` e `botLocation`. Agora, você as usará para gerar novas mensagens no console.

Nos vídeos da aula, você aprendeu como trabalhar com concatenação de strings usando o operador `+` para concatenar strings desta maneira:

```js
let firstName = "John";

console.log("Hello, my name is " + firstName + ".");
// result: "Hello, my name is John."
```

Lembre-se de que você precisa estar atento aos espaços ao concatenar strings com variáveis.

Crie uma variável chamada `botIntroduction`.

Em seguida, use a concatenação de strings com o operador `+` para juntar a string `"My name is "` seguida pela variável `bot` e, em seguida, um ponto (`.`).

Atribua este valor à variável `botIntroduction`.

Então, registre a variável `botIntroduction` no console.

# --hints--

Você deve ter uma variável chamada `botIntroduction`.

```js
assert.isNotNull(botIntroduction);
```

A variável `botIntroduction` deve armazenar o valor de uma string.

```js
assert.isString(botIntroduction);
```

Você deve usar concatenação de strings com o operador `+` para juntar a string `"My name is "` com a variável `bot` seguida por um ponto (`.`). Fique atento aos espaços.

```js
assert.equal(botIntroduction, "My name is teacherBot.");
```

You should have a fourth `console.log()` statement in your code.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 4);
```

A instrução `console` deve gerar a variável `botIntroduction`.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*botIntroduction\s*\);?/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

--fcc-editable-region--

--fcc-editable-region--
```
