---
id: 66adc5919c1853448f119d11
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

Agora parece que o bot quer mudar de apelido.

Usando reatribuição, atribua a string `"awesomeTeacherBot"` à variável `bot`.

# --hints--

Você deve atribuir a string `"awesomeTeacherBot"` à variável `bot`.

```js
assert.strictEqual(bot, "awesomeTeacherBot");
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

--fcc-editable-region--

--fcc-editable-region--
```
