---
id: 66ad8ab945d266383f318cbf
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Agora você deve ver a primeira mensagem do bot no console.

É hora de adicionar uma segunda mensagem do bot.

Adicione outra instrução de `console.log` para emitir a mensagem `"I am excited to talk to you."` no console.

# --hints--

Você deve ter uma instrução de `console.log()` no código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 2);
```

A instrução `console` deve ter a mensagem `"I am excited to talk to you."`. Não se esqueça das aspas.

```js
assert.match(__helpers.removeJSComments(code), /("|')I\s+am\s+excited\s+to\s+talk\s+to\s+you.\1/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
--fcc-editable-region--

--fcc-editable-region--
```
