---
id: 66adc6046acc18453decf577
title: Passo 12
challengeType: 1
dashedName: step-12
---

# --description--

Para ver o novo apelido do bot, você precisará registrar uma nova mensagem no console.

Crie uma variável chamada `newNicknameGreeting`.

Em seguida, use concatenação de strings para juntar a string `"I love my nickname but I wish people would call me "` com a variável `bot`, seguida de um ponto final.

Atribua o resultado à variável `newNicknameGreeting`.

Em seguida, registre o valor de `newNicknameGreeting` no console.

# --hints--

Você deve ter uma variável chamada `newNicknameGreeting`.

```js
assert.isNotNull(newNicknameGreeting);
```

A variável `newNicknameGreeting` deve ser uma string.

```js
assert.isString(newNicknameGreeting);
```

Você deve usar concatenação de strings para juntar a string `"I love my nickname but I wish people would call me "` com a variável `bot`, seguida de um ponto final.

```js
assert.strictEqual(newNicknameGreeting, "I love my nickname but I wish people would call me awesomeTeacherBot.");
```

Você deve ter uma instrução de `console` no código.

```js
assert.lengthOf(__helpers.removeJSComments(code).match(/console\.log(.*)/g), 7);
```

Você deve registrar o valor de `newNicknameGreeting` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*newNicknameGreeting\s*\)/);
```

# --seed--

## --seed-contents--

```js
console.log("Hi there!");
console.log("I am excited to talk to you.");

let bot;
let botLocation;

bot = "teacherBot";
botLocation = "the universe";

console.log("Allow me to introduce myself.");

const botIntroduction = "My name is " + bot + ".";
console.log(botIntroduction);

const botLocationSentence = "I live in " + botLocation + ".";
console.log(botLocationSentence);

bot = "professorBot";

const nicknameIntroduction = "My nickname is " + bot + ".";
console.log(nicknameIntroduction);

bot = "awesomeTeacherBot";

--fcc-editable-region--

--fcc-editable-region--
```
