---
id: 614202874ca576084fca625f
title: Passaggio 16
challengeType: 0
dashedName: step-16
---

# --description--

Ogni ruolo `region` richiede un’etichetta, che aiuta gli utenti di lettori di schermo a capire lo scopo della regione. Un metodo per aggiungere un’etichetta è inserire un elemento intestazione all’interno della regione e poi farvi riferimento con l’attributo `aria-labelledby`.

Aggiungi i seguenti attributi `aria-labelledby` agli elementi `section`:

- `student-info`
- `html-questions`
- `css-questions`

Poi, all’interno di ogni elemento `section`, annida un elemento `h2` con un `id` che corrisponde all’attributo `aria-labelledby` corrispondente. Dai a ogni `h2` un contenuto testuale adatto.

# --hints--

Dovresti dare al primo elemento `section` un attributo `aria-labelledby` con valore `student-info`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.getAttribute('aria-labelledby'), 'student-info');
```

Dovresti dare al secondo elemento `section` un attributo `aria-labelledby` con valore `html-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.getAttribute('aria-labelledby'), 'html-questions');
```

Dovresti dare al terzo elemento `section` un attributo `aria-labelledby` con valore `css-questions`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.getAttribute('aria-labelledby'), 'css-questions');
```

Dovresti annidare un elemento `h2` all’interno del primo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[0]?.querySelectorAll('h2')?.length, 1);
```

Dovresti annidare un elemento `h2` all’interno del secondo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[1]?.querySelectorAll('h2')?.length, 1);
```

Dovresti annidare un elemento `h2` all’interno del terzo elemento `section`.

```js
assert.equal(document.querySelectorAll('section')?.[2]?.querySelectorAll('h2')?.length, 1);
```

Dovresti dare al primo elemento `h2` un attributo `id` con valore `student-info`.

```js
assert.equal(document.querySelectorAll('h2')?.[0]?.id, 'student-info');
```

Dovresti dare al secondo elemento `h2` un attributo `id` con valore `html-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[1]?.id, 'html-questions');
```

Dovresti dare al terzo elemento `h2` un attributo `id` con valore `css-questions`.

```js
assert.equal(document.querySelectorAll('h2')?.[2]?.id, 'css-questions');
```

Dovresti dare al primo elemento `h2` un contenuto testuale adatto. _Suggerimento: avrei scelto `Student Info`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[0]?.textContent.trim());
```

Dovresti dare al secondo elemento `h2` un contenuto testuale adatto. _Suggerimento: avrei scelto `HTML`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[1]?.textContent.trim());
```

Dovresti dare al terzo elemento `h2` un contenuto testuale adatto. _Suggerimento: avrei scelto `CSS`_

```js
assert.isNotEmpty(document.querySelectorAll('h2')[2]?.textContent.trim());
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <title>Accessibility Quiz</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <header>
      <img id="logo" alt="freeCodeCamp" src="https://cdn.freecodecamp.org/platform/universal/fcc_primary.svg">
      <h1>HTML/CSS Quiz</h1>
      <nav>
        <ul>
          <li><a>INFO</a></li>
          <li><a>HTML</a></li>
          <li><a>CSS</a></li>
        </ul>
      </nav>
    </header>
    <main>
--fcc-editable-region--
      <form method="post" action="https://freecodecamp.org/practice-project/accessibility-quiz">
        <section role="region"></section>
        <section role="region"></section>
        <section role="region"></section>
      </form>
--fcc-editable-region--
    </main>
  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}

header {
  width: 100%;
  height: 50px;
  background-color: #1b1b32;
  display: flex;
}

#logo {
  width: max(10rem, 18vw);
  background-color: #0a0a23;
  aspect-ratio: 35 / 4;
  padding: 0.4rem;
}

h1 {
  color: #f1be32;
  font-size: min(5vw, 1.2em);
}

nav {
  width: 50%;
  max-width: 300px;
  height: 50px;
}

nav > ul {
  display: flex;
  justify-content: space-evenly;
}
```
