---
id: 614ccc21ea91ef1736b9b578
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Per questo workshop, ti è stato fornito un codice boilerplate di base.

Inizia questo percorso sull'accessibilità fornendo un attributo `lang` al tuo elemento `html`. Questo aiuterà i lettori di schermo a identificare la lingua della pagina.

Ricorda anche di collegare il tuo foglio di stile alla pagina.

# --hints--

Dovresti dare all'elemento `html` un attributo `lang`. _Suggerimento: puoi usare il valore `en` per l'inglese._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Dovresti avere un elemento `link` all'interno del tuo elemento `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Il tuo elemento `link` dovrebbe avere `rel="stylesheet"` e `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
