---
id: a6e40f1041b06c996f7b2406
title: створіть функцію для пошуку першого елемента
challengeType: 26
dashedName: build-a-first-element-finder
---

# --description--

У цій лабораторній роботі ви створите функцію, яка переглядає масив і повертає перший елемент, що проходить тестову функцію (також відому як «перевірка істинності»).

Функція ітеруватиме по масиву та перевірятиме кожен елемент за допомогою наданої тестової функції. Наприкінці вона поверне перший елемент, для якого тестова функція повертає `true`.

Приклад:

```js
findElement([1, 3, 5, 8], num => num % 2 === 0) // returns 8
findElement([1, 3, 5], num => num % 2 === 0)    // returns undefined
```

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. У вас має бути функція з іменем `findElement`, яка приймає масив і функцію як аргументи.
2. Функція має повертати перший елемент масиву, що проходить перевірку істинності. Це означає, що виклик переданої функції `func` з аргументом елемента `x` вважається пройденим, якщо `func(x)` дорівнює `true`.
3. Якщо жоден елемент не проходить тест, функція має повертати `undefined`.

# --hints--

У вас має бути функція `findElement`.

```js
assert.isFunction(findElement);
```

`findElement([1, 3, 5, 8, 9, 10], function(num) { return num % 2 === 0; })` має повертати `8`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 8, 9, 10], function (num) {
    return num % 2 === 0;
  }),
  8
);
```

`findElement([1, 3, 5, 9], function(num) { return num % 2 === 0; })` має повертати `undefined`.

```js
assert.strictEqual(
  findElement([1, 3, 5, 9], function (num) {
    return num % 2 === 0;
  }),
  undefined
);
```

`findElement([1, 2, 3, 4], function(num) { return num > 2; })` має повертати `3`.

```js
assert.strictEqual(
  findElement([1, 2, 3, 4], function (num) {
    return num > 2;
  }),
  3
);
```

`findElement(["hello", "world", "javascript"], function(str) { return str.length > 5; })` має повертати `"javascript"`.

```js
assert.strictEqual(
  findElement(["hello", "world", "javascript"], function (str) {
    return str.length > 5;
  }),
  "javascript"
);
```

`findElement(["cat", "dog", "bird"], function(str) { return str.length > 10; })` має повертати `undefined`.

```js
assert.strictEqual(
  findElement(["cat", "dog", "bird"], function (str) {
    return str.length > 10;
  }),
  undefined
);
```

`findElement([2, 4, 6, 8], function(num) { return num % 2 === 0; })` має повертати `2`.

```js
assert.strictEqual(
  findElement([2, 4, 6, 8], function (num) {
    return num % 2 === 0;
  }),
  2
);
```

`findElement([], function(num) { return num > 0; })` має повертати `undefined`.

```js
assert.strictEqual(
  findElement([], function (num) {
    return num > 0;
  }),
  undefined
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findElement(arr, func) {
  return arr.filter(func)[0];
}

findElement([1, 2, 3, 4], num => num % 2 === 0);
```
