---
id: 5f33071498eb2472b87ddee4
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

In questo progetto imparerai le basi del CSS (Cascading Style Sheets) costruendo un menu per un caffè. Il CSS è il linguaggio usato per definire lo stile di un documento HTML. Descrive come gli elementi HTML devono essere visualizzati sullo schermo.

Come hai imparato negli ultimi passaggi dell'app Cat Photo, esiste una struttura di base necessaria per iniziare a costruire la tua pagina web. Ogni documento HTML dovrebbe avere una dichiarazione `DOCTYPE` e un elemento `html`. La `DOCTYPE` indica al browser quale versione di HTML è usata nel documento. E l'elemento `html` rappresenta l'elemento radice che contiene tutti gli altri elementi.

```html
<!DOCTYPE html>
<html lang="en">
<!--all other elements go here-->
</html>
```

Aggiungi il tag `<!DOCTYPE html>` e un elemento `html` con un attributo `lang` impostato su `en`.

# --hints--

Dovresti avere la dichiarazione `<!DOCTYPE html>`.

```js
assert(code.match(/<!DOCTYPE html>/i));
```

Dovresti avere un tag di apertura `<html lang="en">` con la lingua impostata su inglese.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Dovresti avere un tag di chiusura `</html>`.

```js
assert(code.match(/<\/html>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--

```
