---
id: 5f331e55dfab7a896e53c3a1
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

Il `title` è uno dei vari elementi che forniscono informazioni aggiuntive non visibili sulla pagina web, ma utili per i motori di ricerca o per come la pagina viene visualizzata.

All'interno dell'elemento `head`, annida un elemento `meta` con un attributo chiamato `charset` impostato sul valore `utf-8` per indicare al browser come codificare i caratteri della pagina.

Ricorda che gli elementi `meta` sono elementi void che non richiedono un tag di chiusura.

# --hints--

Dovresti avere un tag `meta`.

```js
assert.match(code, /<meta.*?\/?>/is);
```

L'elemento `meta` è un elemento void, non dovrebbe avere un tag di chiusura `</meta>`.

```js
assert.notMatch(code, /<\/meta>/i);
```

Il tuo tag `meta` dovrebbe avere un attributo `charset`.

```js
assert.match(code, /<meta\s+charset\s*/i);
```

Il tuo attributo `charset` dovrebbe avere un valore `utf-8`.

```js
assert.match(code, /charset\s*=\s*('|")utf-8\1/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
</html>
```
