---
id: 5f33294a6af5e9188dbdb8f3
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

È il momento di aggiungere del contenuto al menu. Aggiungi un elemento `main` all’interno dell’elemento `body` esistente. Alla fine conterrà informazioni sui prezzi di caffè e dessert offerti dal bar.

# --hints--

Il tuo codice dovrebbe avere un tag di apertura `<main>`.

```js
assert(code.match(/<main>/i));
```

Il tuo codice dovrebbe avere un tag di chiusura `</main>`.

```js
assert(code.match(/<\/main>/i));
```

Non dovresti modificare il tuo elemento `body`. Assicurati di non eliminare accidentalmente il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('body'), 1);
```

Il tuo tag `main` dovrebbe essere all’interno del tag `body`.

```js
const main = document.querySelector('main');
assert(main.parentElement.tagName === 'BODY');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
--fcc-editable-region--
</html>
```
