---
id: 5f3477ae8466a9a3d2cc953c
title: Passaggio 15
challengeType: 0
dashedName: step-15
---

# --description--

Ora che hai il CSS nel file `styles.css`, procedi a rimuovere l'elemento `style` e tutto il suo contenuto. Una volta rimosso, il testo che era centrato tornerà allineato a sinistra.

# --hints--

Non dovresti avere tag `style` nel tuo codice.

```js
assert(!code.match(/style/i));
```

Non dovresti avere selettori CSS nel tuo file HTML.

```js
assert(!editableContents.includes('style'));
assert(!editableContents.includes('text-align'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
--fcc-editable-region--
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
h1, h2, p {
  text-align: center;
}
```
