---
id: 5f3477ae34c1239cafe128be
title: Passaggio 13
challengeType: 0
dashedName: step-13
---

# --description--

Ora hai tre selettori di tipo con lo stesso stile. Puoi aggiungere lo stesso gruppo di stili a molti elementi creando una lista di selettori. Ogni selettore è separato da virgole in questo modo:

```css
selector1, selector2 {
  property: value;
}
```

Elimina i tre selettori di tipo esistenti e sostituiscili con una lista di selettori che centri il testo per gli elementi `h1`, `h2` e `p`.

# --hints--

Dovresti usare un singolo selettore di tipo per tutti e tre gli elementi, `h1, h2, p`. Assicurati di usare quell’ordine.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1, h2, p');
assert(hasSelector);
```

Dovresti avere solo un selettore nel tuo elemento `style`.

```js
const selectors = new __helpers.CSSHelp(document).getCSSRules().filter(x => x.selectorText)
assert(selectors.length === 1);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
      h1 {
        text-align: center;
      }
      h2 {
        text-align: center;
      }
      p {
        text-align: center;
      }
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
