---
id: 5f344f9c805cd193c33d829c
title: Passaggio 11
challengeType: 0
dashedName: step-11
---

# --description--

Puoi definire lo stile di un elemento specificandolo nell'elemento `style` e impostando una proprietà per esso in questo modo:

```css
element {
 property: value;
}
```

Centra il contenuto dell'elemento `h1` impostando la sua proprietà `text-align` al valore `center`.

# --hints--

Dovresti avere un selettore `h1` nel tuo elemento `style`.

```js
const hasSelector = new __helpers.CSSHelp(document).getStyle('h1');
assert(hasSelector);
```

La tua proprietà `text-align` dovrebbe avere un valore `center`.

```js
const hasTextAlign = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['text-align'] === 'center');
assert(hasTextAlign);
```

Il tuo selettore `h1` dovrebbe impostare la proprietà `text-align` su `center`.

```js
const h1TextAlign = new __helpers.CSSHelp(document).getStyle('h1')?.getPropertyValue('text-align');
assert(h1TextAlign === 'center');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
--fcc-editable-region--
    <style>
    </style>
--fcc-editable-region--
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```
