---
id: 5f3477aefa51bfc29327200b
title: Step 14
challengeType: 0
dashedName: step-14
---

# --description--

Hai impostato lo stile di tre elementi scrivendo del CSS all'interno dei tag dell'elemento `style`. Questo approccio funziona, ma poiché ci saranno molti altri stili, è meglio mettere tutti gli stili in un file separato e collegarsi ad esso.

A separate `styles.css` file has been created for you. You can switch between files with the tabs at the top of the editor.

Inizia riscrivendo gli stili che hai creato all'interno del file `styles.css`. Make sure you exclude the opening and closing `style` tags.

# --hints--

Il file `styles.css` dovrebbe avere il selettore di tipo `h1, h2, p`.

```js
assert(editableContents.replace(/[\s\n]*/g, "").match(/(h1|h2|p),(h1|h2|p),(h1|h2|p){/));
```

Il selettore dovrebbe assegnare alla proprietà `text-align` il valore `center`.

```js
assert(editableContents.match(/text-align:\s*center;?/));
```

Il codice non dovrebbe contenere i tag `<style>`.

```js
assert.isFalse(/<\/html>(\n|.)*<\/?\s*style\s*>/.test(code));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
    <style>
      h1, h2, p {
        text-align: center;
      }
    </style>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
      <section>
        <h2>Coffee</h2>
      </section>
    </main>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--

```
