---
id: 5f344fbc22624a2976425065
title: Passaggio 9
challengeType: 0
dashedName: step-9
---

# --description--

Crea un elemento `h2` all'interno dell'elemento `section` e assegnagli il testo `Coffee`.

# --hints--

Dovresti avere un tag di apertura `<h2>`.

```js
assert(code.match(/<h2\s*>/i));
```

Dovresti avere un tag di chiusura `</h2>`.

```js
assert(code.match(/<\/h2\s*>/i));
```

Non dovresti modificare il tuo elemento `section` esistente. Assicurati di non aver eliminato il tag di chiusura.

```js
assert.lengthOf(document.querySelectorAll('section'),1);
```

Il tuo elemento `h2` dovrebbe essere all'interno del tuo elemento `section`.

```js
const h2 = document.querySelector('h2');
assert(h2.parentElement.tagName === 'SECTION');
```

Il tuo elemento `h2` dovrebbe contenere il testo `Coffee`.

```js
const h2 = document.querySelector('h2');
assert(h2.innerText === 'Coffee');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <title>Cafe Menu</title>
  </head>
  <body>
    <main>
      <h1>CAMPER CAFE</h1>
      <p>Est. 2020</p>
--fcc-editable-region--
      <section>
      </section>
--fcc-editable-region--
    </main>
  </body>
</html>
```
