---
id: 5f35e5c44359872a137bd98f
title: Passaggio 28
challengeType: 0
dashedName: step-28
---

# --description--

Dato che il prodotto principale in vendita nel caffè è il caffè, potresti usare un'immagine di chicchi di caffè come sfondo della pagina.

Rimuovi il commento e il suo contenuto all'interno del selettore di tipo `body`. Dopodiché, aggiungi una proprietà `background-image` e imposta il suo valore su `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

# --hints--

Dovresti rimuovere la proprietà `background-color` commentata.

```js
assert(!code.match(/\/\*\s*background-color:\s*burlywood;?\s*\*\//i))
```

Il tuo selettore `body` non dovrebbe contenere commenti.

```js
assert(!code.match(/body\s*{\s*\/\*/i));
```

Dovresti impostare la proprietà `background-image` su `url(https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg)`.

```js
const hasBackground = new __helpers.CSSHelp(document).getCSSRules().some(x => x.style['background-image'] === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`)
assert(hasBackground)
```

Il tuo elemento `body` dovrebbe avere l'immagine di sfondo dei chicchi di caffè.

```js
const bodyBackground = new __helpers.CSSHelp(document).getStyle('body')?.getPropertyValue('background-image');

assert(bodyBackground === `url("https://cdn.freecodecamp.org/curriculum/css-cafe/beans.jpg")`);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Cafe Menu</title>
    <link href="styles.css" rel="stylesheet"/>
  </head>
  <body>
    <div class="menu">
      <main>
        <h1>CAMPER CAFE</h1>
        <p>Est. 2020</p>
        <section>
          <h2>Coffee</h2>
        </section>
      </main>
    </div>
  </body>
</html>
```

```css
--fcc-editable-region--
body {
  /*
  background-color: burlywood;
  */
}
--fcc-editable-region--

h1, h2, p {
  text-align: center;
}

.menu {
  width: 80%;
  background-color: burlywood;
  margin-left: auto;
  margin-right: auto;
}
```
