---
id: 66bba6fff611169359d9d36a
title: Erstellen Sie eine vertrauliche E-Mail-Seite
challengeType: 25
dashedName: build-a-confidential-email-page
demoType: onClick
---

# --description--

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten ein `main`-Element mit einem `id` von `email` haben.
1. Ihr `#email`-Element sollte einen `padding` von `50px`, einen oberen Außenabstand von `50px`, einen `width` von `500px` und einen `border` haben, der `2px` ist.
1. Die Gesamtbreite Ihres `#email`-Elements, einschließlich Innenabstände und Rahmen, sollte `500px` betragen.
1. Sie sollten zwei `div`-Elemente haben, eines mit der ID `confidential` und das andere mit der ID `top-secret`, innerhalb Ihres `#email`-Elements.
1. Ihre `#confidential`- und `#top-secret`-Elemente sollten einen `display` von `inline-block` haben.
1. Ihre `#confidential`- und `#top-secret`-Elemente sollten einen `padding`, einen linken Außenabstand und einen `border` haben.
1. Das `#confidential`-Element sollte den Text `CONFIDENTIAL` enthalten.
1. Das `#top-secret`-Element sollte den Text `TOP SECRET` enthalten.
1. Ihre `#confidential`- und `#top-secret`-Elemente sollten mit einer CSS-Transformation rotiert werden.
1. Sie sollten mindestens drei Absatz-Elemente innerhalb Ihres `#email`-Elements haben.
1. Sie sollten mindestens drei `span`-Elemente mit der Klasse `blurred` innerhalb Ihrer Absatz-Elemente haben.
1. Sie sollten einen `blurred`-Selektor haben, der das Element `3px` mit einem CSS-Filter unscharf stellt.

# --hints--

Sie sollten ein `main`-Element mit einem `id` von `email` haben.

```js
assert.exists(document.querySelector('main#email'));
```

Sie sollten einen `#email`-Selektor haben, der den `padding` seiner Elemente auf allen Seiten auf `50px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('padding'), '50px');
```

Sie sollten einen `#email`-Selektor haben, der den `margin-top` seiner Elemente auf `50px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('margin-top'), '50px');
```

Sie sollten einen `#email`-Selektor haben, der den `width` seiner Elemente auf `500px` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('width'), '500px');
```

Ihr `#email`-Element sollte einen `2px`-Rahmen haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.getPropertyValue('border-width'), '2px');
```

Ihr `#email`-Element sollte einen `box-sizing` von `border-box` haben.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('#email')?.boxSizing, 'border-box');
```

Sie sollten ein `div`-Element mit einem `id` von `confidential` innerhalb Ihres `#email`-Elements haben.

```js
assert.exists(document.querySelector('#email > div#confidential'));
```

Sie sollten ein `div`-Element mit einem `id` von `top-secret` innerhalb Ihres `#email`-Elements haben.

```js
assert.exists(document.querySelector('#email > div#top-secret'));
```

Ihr `#confidential`-Element sollte seinen `display` auf `inline-block` gesetzt haben.

```js
assert.equal(window.getComputedStyle(document.querySelector('#confidential'))?.display, 'inline-block');
```

Ihr `#top-secret`-Element sollte seinen `display` auf `inline-block` gesetzt haben.

```js
assert.equal(window.getComputedStyle(document.querySelector('#top-secret'))?.display, 'inline-block');
```

Ihr `#confidential`-Element sollte auf allen Seiten einen `padding` haben.

```js
const style = window.getComputedStyle(document.querySelector('#confidential'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ihr `#top-secret`-Element sollte auf allen Seiten einen `padding` haben.

```js
const style = window.getComputedStyle(document.querySelector('#top-secret'));
assert.notEqual(style?.paddingTop, '0px');
assert.notEqual(style?.paddingRight, '0px');
assert.notEqual(style?.paddingBottom, '0px');
assert.notEqual(style?.paddingLeft, '0px');
```

Ihr `#confidential`-Element sollte einen linken Außenabstand haben.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.marginLeft, '0px');
```

Ihr `#top-secret`-Element sollte einen linken Außenabstand haben.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.marginLeft, '0px');
```

Ihr `#confidential`-Element sollte einen `border` haben.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.borderWidth, '0px');
```

Ihr `#top-secret`-Element sollte einen `border` haben.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.borderWidth, '0px');
```

Der Text Ihres `#confidential`-Elements sollte `CONFIDENTIAL` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.equal(document.querySelector('#confidential').innerText.toLowerCase(), 'confidential' );
```

Der Text Ihres `#top-secret`-Elements sollte `TOP SECRET` sein. Sie haben entweder den Text ausgelassen oder einen Tippfehler.

```js
assert.equal(document.querySelector('#top-secret').innerText.toLowerCase(), 'top secret' );
```

Ihr `#confidential`-Element sollte eine CSS-`transform` verwenden, um das Element zu rotieren.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#confidential'))?.transform, 'none');
```

Ihr `#top-secret`-Element sollte eine CSS-`transform` verwenden, um das Element zu rotieren.

```js
assert.notEqual(window.getComputedStyle(document.querySelector('#top-secret'))?.transform, 'none');
```

Sie sollten mindestens drei Absatz-Elemente innerhalb Ihres `#email`-Elements haben.

```js
assert.isAtLeast(document.querySelectorAll('#email > p').length, 3);
```

Sie sollten mindestens drei `span`-Elemente mit der Klasse `blurred` innerhalb Ihrer Absätze haben.

```js
assert.isAtLeast(document.querySelectorAll('p > span.blurred').length, 3);
```

Ihre `.blurred`-Elemente sollten nicht leer sein.

```js
const els = document.querySelectorAll('.blurred');
assert.isAtLeast(els.length, 1);
els.forEach(el => assert.isAtLeast(el.innerText.length , 1))
```

Sie sollten einen `.blurred`-Selektor haben, der den `filter` seiner Elemente auf `blur(3px)` setzt.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.blurred')?.filter, 'blur(3px)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Confidential Email</title>

</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Secret Marshmallow Mission</title>
  <link rel="stylesheet" href="styles.css">
</head>

<body>
  <main id="email">
    <div id="confidential">CONFIDENTIAL</div>
    <p>Dear Agent <span class="blurred">S'more,</span></p>
    <p>We have an emergency. The secret formula for our <span class="blurred">Mega Marshmallow</span> has been
      compromised. This formula is what makes our marshmallows the fluffiest and most delicious.</p>
    <p>We suspect that <span class="blurred">Professor Puff</span> is behind this. He has taken the formula to his
      hidden laboratory. Your mission is to infiltrate the lab and secure the formula before it's too late.</p>
    <p>Be sure to keep the lab's location confidential. Any leak of this information could jeopardize the entire
      operation.</p>
    <div id="top-secret">TOP SECRET</div>
  </main>
</body>

</html>
```

```css
#email {
  background-color: SeaShell;
  margin: 50px auto;
  padding: 50px;
  width: 500px;
  border: 2px solid black;
  box-shadow: 5px 3px 3px gray;
  box-sizing: border-box;
}

#confidential {
  display: inline-block;
  margin-left: 100px;
  padding: 10px;
  text-align: center;
  font-size: 30px;
  transform: rotate(-20deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

#top-secret {
  margin-left: 200px;
  display: inline-block;
  padding: 10px;
  text-align: center;
  font-size: 20px;
  transform: rotate(15deg);
  color: red;
  border: 5px solid red;
  font-weight: bold;
}

.blurred {
  filter: blur(3px);
}
```
