---
id: 587d7790367417b2b2512ab0
title: "Verwenden Sie tabindex, um einem Element den Tastaturfokus hinzuzufügen"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMDHW'
forumTopicId: 301027
dashedName: use-tabindex-to-add-keyboard-focus-to-an-element
---

# --description--

Das HTML-Attribut `tabindex` hat drei unterschiedliche Funktionen in Bezug auf den Tastaturfokus eines Elements. Wenn es auf einem Tag steht, zeigt es an, dass das Element fokussierbar ist. Der Wert (eine ganze Zahl, die positiv, negativ oder null sein kann) bestimmt das Verhalten.

Bestimmte Elemente, wie Links und Formularsteuerelemente, erhalten automatisch den Tastaturfokus, wenn ein Benutzer mit der Tabulatortaste durch eine Seite navigiert. Die Reihenfolge entspricht der Reihenfolge der Elemente im HTML-Quelltext. Diese gleiche Funktionalität kann anderen Elementen, wie `div`, `span` und `p`, gegeben werden, indem man ihnen ein `tabindex="0"`-Attribut hinzufügt. Hier ist ein Beispiel:

```html
<div tabindex="0">I need keyboard focus!</div>
```

**Hinweis:** Ein negativer `tabindex`-Wert (typischerweise -1) zeigt an, dass ein Element fokussierbar ist, aber nicht über die Tastatur erreichbar ist. Diese Methode wird normalerweise verwendet, um den Fokus programmatisch auf Inhalte zu setzen (zum Beispiel wenn ein `div` für ein Pop-up-Fenster aktiviert wird) und liegt außerhalb des Gültigkeitsbereichs dieser Aufgaben.

# --instructions--

Camper Cat hat eine neue Umfrage erstellt, um Informationen über seine Nutzer zu sammeln. Er weiß, dass Eingabefelder automatisch den Tastaturfokus erhalten, aber er möchte sicherstellen, dass seine Tastaturnutzer bei den Anweisungen anhalten, während sie mit der Tabulatortaste durch die Elemente navigieren. Fügen Sie dem `p`-Tag ein `tabindex`-Attribut hinzu und setzen Sie dessen Wert auf `0`. Bonus – die Verwendung von `tabindex` ermöglicht auch, dass die CSS-Pseudoklasse `:focus` auf dem `p`-Tag funktioniert.

# --hints--

Ihr Code sollte dem `p`-Tag, das die Formularanweisungen enthält, ein `tabindex`-Attribut hinzufügen.

```js
assert.isNotNull(document.querySelector('p')?.getAttribute('tabindex'));
```

Ihr Code sollte das `tabindex`-Attribut auf dem `p`-Tag auf den Wert 0 setzen.

```js
assert.equal(document.querySelector('p')?.getAttribute('tabindex'), '0');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p>Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<head>
  <style>
  p:focus {
    background-color: yellow;
  }
  </style>
</head>
<body>
  <header>
    <h1>Ninja Survey</h1>
  </header>
  <section>
    <form>


      <p tabindex="0">Instructions: Fill in ALL your information then click <b>Submit</b></p>


      <label for="username">Username:</label>
      <input type="text" id="username" name="username"><br>
      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">9th Life Master</label>
      </fieldset>
      <br>
      <fieldset>
      <legend>Select your favorite weapons:</legend>
      <input id="stars" type="checkbox" name="weapons" value="stars">
      <label for="stars">Throwing Stars</label><br>
      <input id="nunchucks" type="checkbox" name="weapons" value="nunchucks">
      <label for="nunchucks">Nunchucks</label><br>
      <input id="sai" type="checkbox" name="weapons" value="sai">
      <label for="sai">Sai Set</label><br>
      <input id="sword" type="checkbox" name="weapons" value="sword">
      <label for="sword">Sword</label>
      </fieldset>
      <br>
      <input type="submit" name="submit" value="Submit">
    </form><br>
  </section>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
