---
id: 587d778a367417b2b2512aa5
title: Barrierefreiheit von Diagrammen mit dem figure-Element verbessern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cGJMqtE'
forumTopicId: 301015
dashedName: improve-chart-accessibility-with-the-figure-element
---

# --description--

HTML5 führte das `figure`-Element und das zugehörige `figcaption` ein. Zusammen verwendet, umschließen diese Elemente eine visuelle Darstellung (wie ein Bild, Diagramm oder Chart) zusammen mit ihrer Bildunterschrift. Das gemeinsame Umschließen dieser Elemente verbessert die Barrierefreiheit auf zweifache Weise, indem verwandte Inhalte semantisch gruppiert und eine Textalternative bereitgestellt wird, die das `figure` erklärt.

Für Datenvisualisierungen wie Diagramme kann die Bildunterschrift verwendet werden, um für Nutzer mit Sehbeeinträchtigung kurz die Trends oder Schlussfolgerungen zu notieren. Eine weitere Aufgabe behandelt, wie man eine Tabellenversion der Diagrammdaten für Screenreader-Nutzer mit CSS außerhalb des Bildschirms platziert.

Hier ein Beispiel – beachten Sie, dass das `figcaption` innerhalb der `figure`-Tags steht und mit anderen Elementen kombiniert werden kann:

```html
<figure>
  <img src="roundhouseDestruction.jpeg" alt="Photo of Camper Cat executing a roundhouse kick">
  <br>
  <figcaption>
    Master Camper Cat demonstrates proper form of a roundhouse kick.
  </figcaption>
</figure>
```

# --instructions--

Camper Cat arbeitet fleißig an einem gestapelten Balkendiagramm, das die wöchentliche Trainingszeit für Heimlichkeit, Kampf und Waffen zeigt. Helfen Sie ihm, seine Seite besser zu strukturieren, indem Sie das von ihm verwendete `div`-Tag in ein `figure`-Tag ändern und das `p`-Tag, das die Bildunterschrift umgibt, in ein `figcaption`-Tag umwandeln.

# --hints--

Ihr Code sollte ein `figure`-Tag enthalten.

```js
assert.lengthOf(document.querySelectorAll('figure') , 1);
```

Ihr Code sollte ein `figcaption`-Tag enthalten.

```js
assert.lengthOf(document.querySelectorAll('figcaption') , 1);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Ihr Code sollte keine `p`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('p') , 0);
```

Das `figcaption` sollte ein Kind des `figure`-Tags sein.

```js
const figure = document.querySelector('figure');
const children = figure?.querySelectorAll(`:scope ${'figcaption'}`);
assert.lengthOf(children, 1);
```

Ihr `figure`-Element sollte ein schließendes Tag haben.

```js
assert.isTrue(
    code.match(/<\/figure>/g)?.length === code.match(/<figure>/g)?.length
);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>

      <!-- Only change code below this line -->
      <div>
        <!-- Stacked bar chart will go here -->
        <br>
        <p>Breakdown per week of time to spend training in stealth, combat, and weapons.</p>
      </div>
      <!-- Only change code above this line -->

    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section>
      <figure>
        <!-- Stacked bar chart will go here -->
        <br>
        <figcaption>Breakdown per week of time to spend training in stealth, combat, and weapons.</figcaption>
      </figure>
    </section>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
