---
id: 587d7788367417b2b2512aa3
title: Navigation für Screenreader mit dem Landmark-Tag „footer“ erleichtern
challengeType: 0
videoUrl: 'https://scrimba.com/c/crVrDh8'
forumTopicId: 301022
dashedName: make-screen-reader-navigation-easier-with-the-footer-landmark
---

# --description--

Ähnlich wie `header` und `nav` verfügt das `footer`-Element über eine eingebaute Landmark-Funktion, die es assistiven Geräten ermöglicht, schnell zu ihm zu navigieren. Es wird hauptsächlich verwendet, um Urheberrechtsinformationen oder Links zu verwandten Dokumenten zu enthalten, die normalerweise am unteren Rand einer Seite stehen.

# --instructions--

Die Trainingsseite von Camper Cat macht gute Fortschritte. Ändern Sie das `div`, das er verwendet hat, um seine Urheberrechtsinformationen am unteren Rand der Seite einzuschließen, in ein `footer`-Element.

# --hints--

Ihr Code sollte ein `footer`-Tag enthalten.

```js
assert.lengthOf(document.querySelectorAll('footer') ,1);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

Ihr Code sollte ein öffnendes und ein schließendes `footer`-Tag enthalten.

```js
assert.match(code,/<footer>\s*&copy; 2018 Camper Cat\s*<\/footer>/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>


  <div>&copy; 2018 Camper Cat</div>


</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training</h1>
    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>
  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>


  <footer>&copy; 2018 Camper Cat</footer>


</body>
```
