---
id: 587d778e367417b2b2512aab
title: Lesbarkeit mit Text in hohem Kontrast verbessern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cKb3nCq'
forumTopicId: 301017
dashedName: improve-readability-with-high-contrast-text
---

# --description--

Ein geringer Kontrast zwischen Vordergrund- und Hintergrundfarben kann das Lesen von Text erschweren. Ein ausreichender Kontrast verbessert die Lesbarkeit Ihres Inhalts, aber was genau bedeutet „ausreichend“?

Die Web Content Accessibility Guidelines (WCAG) empfehlen für normalen Text ein Kontrastverhältnis von mindestens 4,5 zu 1. Das Verhältnis wird berechnet, indem die relative Leuchtdichte zweier Farben verglichen wird. Es reicht von 1:1 für dieselbe Farbe, also ohne Kontrast, bis zu 21:1 für Weiß auf Schwarz, dem stärksten Kontrast. Es gibt viele Online-Tools zur Kontrastprüfung, die dieses Verhältnis für Sie berechnen.

# --instructions--

Camper Cats Wahl von hellgrauem Text auf weißem Hintergrund für seinen aktuellen Blogbeitrag hat ein Kontrastverhältnis von 1,5:1, was das Lesen erschwert. Ändern Sie die `color` des Texts von dem aktuellen Grau (`#D3D3D3`) zu einem dunkleren Grau (`#636363`), um das Kontrastverhältnis auf 6:1 zu verbessern.

# --hints--

Ihr Code sollte die `color` des Texts für die `body` in das dunklere Grau ändern.

```js
const body = document.querySelector('body');
const bodyColor = window.getComputedStyle(body).color; 
assert(bodyColor == 'rgb(99, 99, 99)');
```

Ihr Code sollte die `background-color` für die `body` nicht ändern.

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body).backgroundColor; 
assert.equal(backgroundColor , 'rgb(255, 255, 255)');
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: #D3D3D3;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: #636363;
    background-color: #FFF;
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
