---
id: 587d7788367417b2b2512aa2
title: Navigation für Screenreader mit dem nav-Landmark-Element erleichtern
challengeType: 0
videoUrl: 'https://scrimba.com/c/czVwWSv'
forumTopicId: 301024
dashedName: make-screen-reader-navigation-easier-with-the-nav-landmark
---

# --description--

Das `nav`-Element ist ein weiteres HTML5-Element mit der eingebetteten Landmark-Funktion für eine einfache Navigation mit Screenreadern. Dieses Tag soll die Hauptnavigationslinks auf Ihrer Seite umschließen.

Wenn sich am unteren Rand der Seite wiederholte Seitenlinks befinden, ist es nicht notwendig, diese ebenfalls mit einem `nav`-Tag auszuzeichnen. Die Verwendung eines `footer` (im nächsten Challenge behandelt) ist ausreichend.

# --instructions--

Camper Cat hat Navigationslinks oben auf seiner Trainingsseite eingefügt, diese aber in einem `div` umschlossen. Ändern Sie das `div` in ein `nav`-Tag, um die Barrierefreiheit seiner Seite zu verbessern.

# --hints--

Ihr Code sollte genau ein `nav`-Tag enthalten.

```js
assert.lengthOf(document.querySelectorAll('nav') , 1);
```

Ihre `nav`-Tags sollten das `ul` und dessen Listenelemente umschließen.

```js
const nav = document.querySelector('nav');
const children = nav?.querySelectorAll(`:scope ${'ul'}`);
assert.lengthOf(children,1);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Ihr `nav`-Element sollte ein schließendes Tag haben.

```js
assert.isTrue(code.match(/<\/nav>/g)?.length === code.match(/<nav>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <div>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </div>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Training with Camper Cat</h1>

    <nav>
      <ul>
        <li><a href="#stealth">Stealth &amp; Agility</a></li>
        <li><a href="#combat">Combat</a></li>
        <li><a href="#weapons">Weapons</a></li>
      </ul>
    </nav>

  </header>
  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
