---
id: 587d774c367417b2b2512a9d
title: "Wissen, wann Alt-Text leer bleiben sollte"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cM9P4t2'
forumTopicId: 301019
dashedName: know-when-alt-text-should-be-left-blank
---

# --description--

In der letzten Aufgabe haben Sie gelernt, dass das Hinzufügen eines `alt`-Attributs bei der Verwendung von `img`-Tags verpflichtend ist. Manchmal werden Bilder jedoch mit einer Bildunterschrift gruppiert, die sie bereits beschreibt, oder sie dienen nur der Dekoration. In diesen Fällen kann der `alt`-Text redundant oder unnötig erscheinen.

Wenn ein Bild bereits durch Textinhalt erklärt wird oder keine Bedeutung für eine Seite hinzufügt, benötigt das `img` dennoch ein `alt`-Attribut, das aber auf einen leeren String gesetzt werden kann. Hier ist ein Beispiel:

```html
<img src="visualDecoration.jpeg" alt="">
```

Hintergrundbilder fallen normalerweise ebenfalls unter die Bezeichnung „dekorativ“. Sie werden jedoch typischerweise mit CSS-Regeln angewendet und sind daher kein Teil der Auszeichnung, die Screenreader verarbeiten.

**Hinweis:** Bei Bildern mit einer Bildunterschrift möchten Sie möglicherweise trotzdem `alt`-Text hinzufügen, da dies Suchmaschinen hilft, den Inhalt des Bildes zu katalogisieren.

# --instructions--

Camper Cat hat eine Grundstruktur für den Blogbereich seiner Website programmiert. Er plant, zwischen seinen zwei Artikeln mit einem dekorativen Bild eines Samuraischwerts eine visuelle Trennung einzufügen. Fügen Sie dem `img`-Tag ein `alt`-Attribut hinzu und setzen Sie es auf einen leeren String. (Beachten Sie, dass das Bild-`src` nicht mit einer tatsächlichen Datei verknüpft ist – machen Sie sich keine Sorgen, dass keine Schwerter angezeigt werden.)

# --hints--

Ihr `img`-Tag sollte ein `alt`-Attribut haben.

```js
assert.isTrue(document.querySelector('img')?.hasAttribute('alt'));
```

Das `alt`-Attribut sollte auf einen leeren String gesetzt werden.

```js
assert.isEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<article>
  <h2>Defeating your Foe: the Red Dot is Ours!</h2>
  <p>To Come...</p>
</article>

<img src="samuraiSwords.jpeg" alt="">

<article>
  <h2>Is Chuck Norris a Cat Person?</h2>
  <p>To Come...</p>
</article>
```
