---
id: 587d778b367417b2b2512aa8
title: Einen barrierefreien Datumsauswahl-Dialog hinzufügen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cR3bRbCV'
forumTopicId: 301008
dashedName: add-an-accessible-date-picker
---

# --description--

Formulare enthalten oft das `input`-Feld, mit dem verschiedene Formularsteuerelemente erstellt werden können. Das `type`-Attribut dieses Elements gibt an, welche Art von `input`-Element erstellt wird.

Sie haben möglicherweise die `text`- und `submit`-Eingabetypen in früheren Aufgaben bemerkt, und HTML5 hat eine Option eingeführt, um ein `date`-Feld anzugeben. Je nach Browserunterstützung erscheint ein Datumsauswahl-Dialog im `input`-Feld, wenn es den Fokus hat, was das Ausfüllen eines Formulars für alle Benutzer erleichtert.

Für ältere Browser wird der Typ standardmäßig auf `text` gesetzt, daher ist es hilfreich, den Benutzern das erwartete Datumsformat im `label`- oder `placeholder`-Text anzuzeigen, nur für den Fall.

Hier ist ein Beispiel:

```html
<label for="input1">Enter a date:</label>
<input type="date" id="input1" name="input1">
```

# --instructions--

Camper Cat richtet ein Mortal-Kombat-Turnier ein und möchte seine Teilnehmer fragen, welches Datum am besten passt. Fügen Sie ein `input`-Tag mit einem `type`-Attribut von `date`, einem `id`-Attribut von `pickdate` und einem `name`-Attribut von `date` hinzu.

# --hints--

Ihr Code sollte ein `input`-Tag für das Datumsauswahlfeld hinzufügen.

```js
assert.lengthOf(document.querySelectorAll('input'), 2);
```

Ihr `input`-Tag sollte ein `type`-Attribut mit dem Wert `date` haben.

```js
assert.equal(document.querySelector('input')?.getAttribute('type'), 'date');
```

Ihr `input`-Tag sollte ein `id`-Attribut mit dem Wert `pickdate` haben.

```js
assert.equal(document.querySelector('input')?.getAttribute('id'),'pickdate');
```

Ihr `input`-Tag sollte ein `name`-Attribut mit dem Wert `date` haben.

```js
assert.equal(document.querySelector('input')?.getAttribute('name'), 'date');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>

        <!-- Only change code below this line -->



        <!-- Only change code above this line -->

        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <main>
    <section>
      <h2>Mortal Kombat Tournament Survey</h2>
      <form>
        <p>Tell us the best date for the competition</p>
        <label for="pickdate">Preferred Date:</label>
        <input type="date" id="pickdate" name="date">
        <input type="submit" name="submit" value="Submit">
      </form>
    </section>
  </main>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
