---
id: 587d778b367417b2b2512aa7
title: Radio-Buttons in einem Fieldset-Element für bessere Barrierefreiheit umschließen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cVJVefw'
forumTopicId: 301030
dashedName: wrap-radio-buttons-in-a-fieldset-element-for-better-accessibility
---

# --description--

Das nächste Formular-Thema behandelt die Barrierefreiheit von Radio-Buttons. Jede Auswahl erhält ein `label` mit einem `for`-Attribut, das mit dem `id` des entsprechenden Elements verknüpft ist, wie in der letzten Aufgabe behandelt. Da Radio-Buttons oft in einer Gruppe vorkommen, bei der der Benutzer eine auswählen muss, gibt es eine Möglichkeit, semantisch zu zeigen, dass die Auswahl zu einer Gruppe gehört.

Das `fieldset`-Tag umschließt die gesamte Gruppe von Radio-Buttons, um dies zu erreichen. Es verwendet oft ein `legend`-Tag, um eine Beschreibung für die Gruppe bereitzustellen, die Screenreader für jede Auswahl im `fieldset`-Element vorlesen.

Der `fieldset`-Wrapper und das `legend`-Tag sind nicht notwendig, wenn die Auswahl selbsterklärend ist, wie bei einer Geschlechtsauswahl. Die Verwendung eines `label` mit dem `for`-Attribut für jeden Radio-Button ist ausreichend.

Hier ist ein Beispiel:

```html
<form>
  <fieldset>
    <legend>Choose one of these three items:</legend>
    <input id="one" type="radio" name="items" value="one">
    <label for="one">Choice One</label><br>
    <input id="two" type="radio" name="items" value="two">
    <label for="two">Choice Two</label><br>
    <input id="three" type="radio" name="items" value="three">
    <label for="three">Choice Three</label>
  </fieldset>
</form>
```

# --instructions--

Camper Cat möchte Informationen über das Ninja-Level seiner Nutzer erhalten, wenn sie sich für seine E-Mail-Liste anmelden. Er hat eine Gruppe von Radio-Buttons hinzugefügt und aus unserer letzten Lektion gelernt, für jede Auswahl `label`-Tags mit `for`-Attributen zu verwenden. Weiter so, Camper Cat! Sein Code braucht jedoch noch etwas Hilfe. Ändern Sie das `div`-Tag, das die Radio-Buttons umgibt, in ein `fieldset`-Tag, und ändern Sie das `p`-Tag darin in ein `legend`.

# --hints--

Ihr Code sollte ein `fieldset`-Tag um die Gruppe der Radio-Buttons haben.

```js
assert.lengthOf(document.querySelectorAll('fieldset'),1);
```

Das `fieldset`-Element sollte ein schließendes Tag haben.

```js
assert.isTrue(code.match(/<\/fieldset>/g)?.length === code.match(/<fieldset>/g)?.length);
```

Ihr Code sollte ein `legend`-Tag um den Text haben, der fragt, welches Ninja-Level ein Benutzer hat.

```js
assert.lengthOf(document.querySelectorAll('legend') ,1);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div') ,0);
```

Ihr Code sollte kein `p`-Tag mehr um den Text haben, der fragt, welches Ninja-Level ein Benutzer hat.

```js
assert.lengthOf(document.querySelectorAll('p') ,4);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">


      <!-- Only change code below this line -->
      <div>
        <p>What level ninja are you?</p>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </div>
      <!-- Only change code above this line -->


      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <section>
    <form>
      <p>Sign up to receive Camper Cat's blog posts by email here!</p>
      <label for="email">Email:</label>
      <input type="text" id="email" name="email">

      <fieldset>
        <legend>What level ninja are you?</legend>
        <input id="newbie" type="radio" name="levels" value="newbie">
        <label for="newbie">Newbie Kitten</label><br>
        <input id="intermediate" type="radio" name="levels" value="intermediate">
        <label for="intermediate">Developing Student</label><br>
        <input id="master" type="radio" name="levels" value="master">
        <label for="master">Master</label>
      </fieldset>

      <input type="submit" name="submit" value="Submit">
    </form>
  </section>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>
  <img src="samuraiSwords.jpeg" alt="">
  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
