---
id: 587d7787367417b2b2512aa1
title: Navigation für Screenreader mit dem header-Landmark erleichtern
challengeType: 0
videoUrl: 'https://scrimba.com/c/cB76vtv'
forumTopicId: 301023
dashedName: make-screen-reader-navigation-easier-with-the-header-landmark
---

# --description--

Das nächste HTML5-Element, das semantische Bedeutung hinzufügt und die Barrierefreiheit verbessert, ist das `header`-Tag. Es wird verwendet, um einleitende Informationen oder Navigationslinks für sein übergeordnetes Tag einzuschließen und eignet sich gut für Inhalte, die oben auf mehreren Seiten wiederholt werden.

`header` teilt die eingebettete Landmark-Funktion, die Sie bei `main` gesehen haben, und ermöglicht es unterstützenden Technologien, schnell zu diesem Inhalt zu navigieren.

**Hinweis:** Das `header` ist für die Verwendung im `body`-Tag Ihres HTML-Dokuments gedacht. Es unterscheidet sich vom `head`-Element, das den Titel der Seite, Meta-Informationen usw. enthält.

# --instructions--

Camper Cat schreibt großartige Artikel über Ninja-Training und möchte eine Seite dafür auf seiner Website hinzufügen. Ändern Sie das obere `div`, das derzeit das `h1` enthält, in ein `header`-Tag.

# --hints--

Ihr Code sollte ein `header`-Tag enthalten.

```js
assert.lengthOf(document.querySelectorAll('header'),1);
```

Ihre `header`-Tags sollten das `h1` umschließen.

```js
const header = document.querySelector('header');
const children = header?.querySelectorAll(`:scope ${'h1'}`);
assert.lengthOf(children , 1);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div') , 0);
```

Ihr `header`-Element sollte ein schließendes Tag haben.

```js
assert.isTrue(code.match(/<\/header>/g)?.length === code.match(/<header>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>

  <div>
    <h1>Training with Camper Cat</h1>
  </div>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```

# --solutions--

```html
<body>

  <header>
    <h1>Training with Camper Cat</h1>
  </header>


  <main>
    <section id="stealth">
      <h2>Stealth &amp; Agility Training</h2>
      <article><h3>Climb foliage quickly using a minimum spanning tree approach</h3></article>
      <article><h3>No training is NP-complete without parkour</h3></article>
    </section>
    <section id="combat">
      <h2>Combat Training</h2>
      <article><h3>Dispatch multiple enemies with multithreaded tactics</h3></article>
      <article><h3>Goodbye world: 5 proven ways to knock out an opponent</h3></article>
    </section>
    <section id="weapons">
      <h2>Weapons Training</h2>
      <article><h3>Swords: the best tool to literally divide and conquer</h3></article>
      <article><h3>Breadth-first or depth-first in multi-weapon training?</h3></article>
    </section>
  </main>
</body>
```
