---
id: 587d7790367417b2b2512aaf
title: Links mit HTML-Zugriffstasten navigierbar machen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cQvmaTp'
forumTopicId: 301021
dashedName: make-links-navigable-with-html-access-keys
---

# --description--

HTML bietet das `accesskey`-Attribut, um eine Tastenkombination festzulegen, mit der ein Element aktiviert oder in den Fokus gebracht wird. Das Hinzufügen eines `accesskey`-Attributs kann die Navigation für Nutzer, die ausschließlich mit der Tastatur arbeiten, effizienter machen.

HTML5 erlaubt die Verwendung dieses Attributs bei jedem Element, aber besonders nützlich ist es bei interaktiven Elementen. Dazu gehören Links, Buttons und Formularsteuerelemente.

Hier ein Beispiel:

```html
<button accesskey="b">Important Button</button>
```

# --instructions--

Camper Cat möchte, dass die Links bei den beiden Blogartikel-Titeln Tastenkombinationen haben, damit die Nutzer seiner Website schnell zur vollständigen Geschichte navigieren können. Fügen Sie beiden Links ein `accesskey`-Attribut hinzu und setzen Sie das erste auf `g` (für Garfield) und das zweite auf `c` (für Chuck Norris).

# --hints--

Ihr Code sollte dem `accesskey`-Tag ein `a`-Attribut mit dem `id` von `first` hinzufügen.

```js
assert.isTrue(document.querySelector('#first')?.hasAttribute('accesskey'));
```

Ihr Code sollte dem `accesskey`-Tag ein `a`-Attribut mit dem `id` von `second` hinzufügen.

```js
assert.isTrue(document.querySelector('#second')?.hasAttribute('accesskey'));
```

Ihr Code sollte das `accesskey`-Attribut im `a`-Tag mit dem `id` von `first` auf `g` setzen. Beachten Sie, dass die Groß- und Kleinschreibung wichtig ist.

```js
assert.equal(document.querySelector('#first')?.getAttribute('accesskey'), 'g');
```

Ihr Code sollte das `accesskey`-Attribut im `a`-Tag mit dem `id` von `second` auf `c` setzen. Beachten Sie, dass die Groß- und Kleinschreibung wichtig ist.

```js
assert.equal(document.querySelector('#second')?.getAttribute('accesskey'), 'c');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>


    <h2><a id="first" accesskey="g" href="#">The Garfield Files: Lasagna as Training Fuel?</a></h2>


    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>
  <article>


    <h2><a id="second" accesskey="c" href="#">Is Chuck Norris a Cat Person?</a></h2>


    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
