---
id: 587d774e367417b2b2512aa0
title: Inhalt im article-Element umschließen
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp79S3'
forumTopicId: 301029
dashedName: wrap-content-in-the-article-element
---

# --description--

`article` ist eines der neuen HTML5-Elemente, die Ihrer Auszeichnung eine semantische Bedeutung hinzufügen. `article` ist ein Abschnittselement und wird verwendet, um unabhängige, eigenständige Inhalte einzuschließen. Der Tag eignet sich gut für Blogeinträge, Forenbeiträge oder Nachrichtenartikel.

Ob Inhalte eigenständig stehen können, ist meist eine Einschätzungssache, aber Sie können ein paar einfache Tests anwenden. Fragen Sie sich, ob der Inhalt auch ohne den umgebenden Kontext noch Sinn ergeben würde. Ebenso, ob der Text in einem RSS-Feed bestehen würde.

Denken Sie daran, dass Menschen, die unterstützende Technologien verwenden, auf eine organisierte, semantisch sinnvolle Auszeichnung angewiesen sind, um Ihre Arbeit besser zu verstehen.

**Hinweis:** Das `section`-Element ist ebenfalls neu mit HTML5 und hat eine leicht andere semantische Bedeutung als `article`. Ein `article` ist für eigenständige Inhalte gedacht, und ein `section` dient zum Gruppieren thematisch verwandter Inhalte. Sie können bei Bedarf ineinander verwendet werden. Zum Beispiel, wenn ein Buch das `article` ist, dann ist jedes Kapitel ein `section`. Wenn keine Beziehung zwischen Inhaltsgruppen besteht, verwenden Sie ein `div`.

`<div>` – gruppiert Inhalte  
`<section>` – gruppiert verwandte Inhalte  
`<article>` – gruppiert unabhängige, eigenständige Inhalte  

# --instructions--

Camper Cat hat `article`-Tags verwendet, um die Beiträge auf seiner Blogseite einzuschließen, aber er hat vergessen, sie um den obersten Beitrag zu setzen. Ändern Sie den `div`-Tag so, dass stattdessen ein `article`-Tag verwendet wird.

# --hints--

Ihr Code sollte drei `article`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('article') ,3);
```

Ihr Code sollte keine `div`-Tags enthalten.

```js
assert.lengthOf(document.querySelectorAll('div'), 0);
```

# --seed--

## --seed-contents--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <div>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </div>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```

# --solutions--

```html
<h1>Deep Thoughts with Master Camper Cat</h1>
<main>
  <article>
    <h2>The Garfield Files: Lasagna as Training Fuel?</h2>
    <p>The internet is littered with varying opinions on nutritional paradigms, from catnip paleo to hairball cleanses. But let's turn our attention to an often overlooked fitness fuel, and examine the protein-carb-NOM trifecta that is lasagna...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near...</p>
  </article>

  <img src="samuraiSwords.jpeg" alt="">

  <article>
    <h2>Is Chuck Norris a Cat Person?</h2>
    <p>Chuck Norris is widely regarded as the premier martial artist on the planet, and it's a complete coincidence that anyone who disagrees with this fact mysteriously disappears soon after. But the real question is, is he a cat person?...</p>
  </article>
</main>
```
