---
id: 587d778c367417b2b2512aa9
title: Zeiten mit dem HTML5-datetime-Attribut standardisieren
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMgtz'
forumTopicId: 301025
dashedName: standardize-times-with-the-html5-datetime-attribute
---

# --description--

Im Anschluss an das Thema Datum hat HTML5 auch das `time`-Element zusammen mit einem `datetime`-Attribut eingeführt, um Zeiten zu standardisieren. Das `time`-Element ist ein Inline-Element, das ein Datum oder eine Uhrzeit auf einer Seite umschließen kann. Ein `datetime`-Attribut enthält ein gültiges Format dieses Datums. Dies ist der Wert, auf den assistive Geräte zugreifen. Es hilft, Verwirrung zu vermeiden, indem eine standardisierte Version einer Zeit angegeben wird, selbst wenn sie im Text informell oder umgangssprachlich geschrieben ist.

Hier ein Beispiel:

```html
<p>Master Camper Cat officiated the cage match between Goro and Scorpion <time datetime="2013-02-13">last Wednesday</time>, which ended in a draw.</p>
```

# --instructions--

Die Umfrageergebnisse von Camper Cat zu Mortal Kombat sind da! Umschließen Sie den Text `time` mit einem `Thursday, September 15<sup>th</sup>`-Tag und fügen Sie diesem ein `datetime`-Attribut hinzu, das auf `2016-09-15` gesetzt ist.

# --before-all--

```js
const getTimeElement = () => {
  const pElement = [...document.querySelectorAll("article > p")]
  .filter(x => x?.textContent?.includes("Thank you to everyone for responding to Master Camper Cat's survey."));
  
  return pElement[0] ? pElement[0].querySelector("time") : null;
};

const getDatetimeAttr = () => {
  const timeElement = getTimeElement();
  return timeElement?.getAttribute("datetime");
};
```

# --hints--

Ihr Code sollte ein `p`-Element enthalten, das den Text `Thank you to everyone for responding to Master Camper Cat's survey.` einschließt, und ein `time`-Element enthalten.

```js
assert.exists(getTimeElement());
```

Ihre hinzugefügten `time`-Tags sollten den Text `Thursday, September 15<sup>th</sup>` umschließen.

```js
assert.strictEqual(getTimeElement()?.innerHTML?.trim(), 'Thursday, September 15<sup>th</sup>');
```

Ihr hinzugefügtes `time`-Tag sollte ein `datetime`-Attribut haben, das nicht leer ist.

```js
assert(getDatetimeAttr()?.length != 0);
```

Ihr hinzugefügtes `datetime`-Attribut sollte auf den Wert `2016-09-15` gesetzt sein.

```js
assert.equal(getDatetimeAttr() , '2016-09-15');
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <article>
    <h2>Mortal Kombat Tournament Survey Results</h2>

    <!-- Only change code below this line -->

    <p>Thank you to everyone for responding to Master Camper Cat's survey. The best day to host the vaunted Mortal Kombat tournament is Thursday, September 15<sup>th</sup>. May the best ninja win!</p>

    <!-- Only change code above this line -->

    <section>
      <h3>Comments:</h3>
      <article>
        <p>Posted by: Sub-Zero on <time datetime="2016-08-13T20:01Z">August 13<sup>th</sup></time></p>
        <p>Johnny Cage better be there, I'll finish him!</p>
      </article>
      <article>
        <p>Posted by: Doge on <time datetime="2016-08-15T08:12Z">August 15<sup>th</sup></time></p>
        <p>Wow, much combat, so mortal.</p>
      </article>
      <article>
        <p>Posted by: The Grim Reaper on <time datetime="2016-08-16T00:00Z">August 16<sup>th</sup></time></p>
        <p>Looks like I'll be busy that day.</p>
      </article>
    </section>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Tournaments</h1>
  </header>
  <article>
    <h2>Mortal Kombat Tournament Survey Results</h2>

    <p>Thank you to everyone for responding to Master Camper Cat's survey. The best day to host the vaunted Mortal Kombat tournament is <time datetime="2016-09-15">Thursday, September 15<sup>th</sup></time>. May the best ninja win!</p>

    <section>
      <h3>Comments:</h3>
      <article>
        <p>Posted by: Sub-Zero on <time datetime="2016-08-13T20:01Z">August 13<sup>th</sup></time></p>
        <p>Johnny Cage better be there, I'll finish him!</p>
      </article>
      <article>
        <p>Posted by: Doge on <time datetime="2016-08-15T08:12Z">August 15<sup>th</sup></time></p>
        <p>Wow, much combat, so mortal.</p>
      </article>
      <article>
        <p>Posted by: The Grim Reaper on <time datetime="2016-08-16T00:00Z">August 16<sup>th</sup></time></p>
        <p>Looks like I'll be busy that day.</p>
      </article>
    </section>
  </article>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
```
