---
id: 587d778f367417b2b2512aac
title: Farbfehlsichtigkeit vermeiden durch ausreichendes Kontrastverhältnis
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzMEUw'
forumTopicId: 301012
dashedName: avoid-colorblindness-issues-by-using-sufficient-contrast
---

# --description--

Farbe ist ein großer Teil des visuellen Designs, aber ihre Verwendung bringt zwei Barrierefreiheitsprobleme mit sich. Erstens sollte Farbe nicht als alleiniges Mittel verwendet werden, um wichtige Informationen zu vermitteln, da Screenreader-Nutzende sie nicht wahrnehmen. Zweitens benötigen Vordergrund- und Hintergrundfarben einen ausreichenden Kontrast, damit farbenblinde Nutzende sie unterscheiden können.

Frühere Aufgaben behandelten Textalternativen, um das erste Problem zu lösen. Die letzte Aufgabe stellte Werkzeuge zur Kontrastprüfung vor, die beim zweiten Problem helfen. Das von den WCAG empfohlene Kontrastverhältnis von 4,5:1 gilt sowohl für Farbkombinationen als auch für Graustufen.

Farbenblinde Nutzende haben Schwierigkeiten, manche Farben voneinander zu unterscheiden – meist im Farbton, manchmal auch in der Helligkeit. Sie erinnern sich vielleicht, dass das Kontrastverhältnis anhand der relativen Luminanz- (oder Helligkeits-) Werte der Vordergrund- und Hintergrundfarben berechnet wird.

In der Praxis kann das Kontrastverhältnis von 4,5:1 erreicht werden, indem man die dunklere Farbe abdunkelt (Schwarz hinzufügt) und die hellere Farbe aufhellt (Weiß hinzufügt). Dunklere Farbtöne im Farbrad sind Blau-, Violett-, Magenta- und Rottöne, während hellere getönte Farben Orange, Gelb, Grün und Blaugrün sind.

# --instructions--

Camper Cat experimentiert mit der Verwendung von Farbe für den Text und Hintergrund seines Blogs, aber seine aktuelle Kombination aus einem grünlichen `background-color` mit kastanienbraunem Text `color` hat ein Kontrastverhältnis von 2,5:1. Sie können die Helligkeit der Farben leicht anpassen, da er sie mit der CSS-Eigenschaft `hsl()` (steht für Farbton, Sättigung, Helligkeit) deklariert hat, indem Sie das dritte Argument ändern. Erhöhen Sie den `background-color`-Helligkeitswert von 35 % auf 55 % und verringern Sie den `color`-Helligkeitswert von 20 % auf 15 %. Dadurch verbessert sich der Kontrast auf 5,9:1.

# --hints--

Ihr Code sollte nur den Helligkeitswert für die Text-`color`-Eigenschaft auf 15 % ändern.

```js
assert.match(code ,/color:\s*?hsl\(0,\s*?55%,\s*?15%\)/gi);
```

Ihr Code sollte nur den Helligkeitswert für die `background-color`-Eigenschaft auf 55 % ändern.

```js
assert.match(code ,/background-color:\s*?hsl\(120,\s*?25%,\s*?55%\)/gi);
```

# --seed--

## --seed-contents--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 20%);
    background-color: hsl(120, 25%, 35%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```

# --solutions--

```html
<head>
  <style>
  body {
    color: hsl(0, 55%, 15%);
    background-color: hsl(120, 25%, 55%);
  }
  </style>
</head>
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>A Word on the Recent Catnip Doping Scandal</h2>
    <p>The influence that catnip has on feline behavior is well-documented, and its use as an herbal supplement in competitive ninja circles remains controversial. Once again, the debate to ban the substance is brought to the public's attention after the high-profile win of Kittytron, a long-time proponent and user of the green stuff, at the Claw of Fury tournament.</p>
    <p>As I've stated in the past, I firmly believe a true ninja's skills must come from within, with no external influences. My own catnip use shall continue as purely recreational.</p>
  </article>
</body>
```
