---
id: 587d7790367417b2b2512ab1
title: "Verwenden Sie tabindex, um die Reihenfolge des Tastaturfokus für mehrere Elemente festzulegen"
challengeType: 0
videoUrl: 'https://scrimba.com/c/cmzRRcb'
forumTopicId: 301028
dashedName: use-tabindex-to-specify-the-order-of-keyboard-focus-for-several-elements
---

# --description--

Das `tabindex`-Attribut legt auch die genaue Tabulatorreihenfolge der Elemente fest. Dies wird erreicht, wenn der Wert des Attributs auf eine positive Zahl von 1 oder höher gesetzt wird.

Das Setzen eines `tabindex="1"` bringt den Tastaturfokus zuerst auf dieses Element. Dann wird die Folge der angegebenen `tabindex`-Werte (2, 3 usw.) durchlaufen, bevor zu den Standard- und `tabindex="0"`-Elementen gewechselt wird.

Es ist wichtig zu beachten, dass bei dieser Art der Tabulatorreihenfolge die Standardreihenfolge (die den HTML-Quellcode verwendet) überschrieben wird. Dies kann Benutzer verwirren, die erwarten, die Navigation von oben auf der Seite zu starten. Diese Technik kann in bestimmten Situationen notwendig sein, aber im Hinblick auf die Barrierefreiheit sollten Sie vorsichtig sein, bevor Sie sie anwenden.

Hier ist ein Beispiel:

```html
<div tabindex="1">I get keyboard focus, and I get it first!</div>
```

```html
<div tabindex="2">I get keyboard focus, and I get it second!</div>
```

# --instructions--

Camper Cat hat auf seiner Seite mit inspirierenden Zitaten ein Suchfeld, das er mit CSS in der oberen rechten Ecke positionieren möchte. Er möchte, dass die Such-`input`- und die Absende-`input`-Formularsteuerelemente die ersten beiden Elemente in der Tabulatorreihenfolge sind. Fügen Sie dem `tabindex` `1` ein `search`-Attribut mit dem Wert `input` hinzu und dem `tabindex` `2` ein `submit`-Attribut mit dem Wert `input`.

Außerdem ist zu beachten, dass einige Browser Sie in der Mitte Ihrer Tabulatorreihenfolge platzieren können, wenn ein Element angeklickt wird. Ein Element wurde der Seite hinzugefügt, das sicherstellt, dass Sie immer am Anfang Ihrer Tabulatorreihenfolge starten.

# --hints--

Ihr Code sollte dem `tabindex` `search`-Tag ein `input`-Attribut hinzufügen.

```js
assert.isNotNull(document.querySelector('#search')?.getAttribute('tabindex'));
```

Ihr Code sollte dem `tabindex` `submit`-Tag ein `input`-Attribut hinzufügen.

```js
assert.isNotNull(document.querySelector('#submit')?.getAttribute('tabindex'));
```

Ihr Code sollte das `tabindex`-Attribut im `search` `input`-Tag auf den Wert 1 setzen.

```js
assert.equal(document.querySelector('#search')?.getAttribute('tabindex'), '1');
```

Ihr Code sollte das `tabindex`-Attribut im `submit` `input`-Tag auf den Wert 2 setzen.

```js
assert.equal(document.querySelector('#submit')?.getAttribute('tabindex'), '2');
```

# --seed--

## --seed-contents--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input type="search" name="search" id="search">
    <input type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```

# --solutions--

```html
<body>
  <div tabindex="1" class="overlay"></div>
  <header>
    <h1>Even Deeper Thoughts with Master Camper Cat</h1>
    <nav>
      <ul>
        <li><a href="">Home</a></li>
        <li><a href="">Blog</a></li>
        <li><a href="">Training</a></li>
      </ul>
    </nav>
  </header>
  <form>
    <label for="search">Search:</label>


    <input tabindex="1" type="search" name="search" id="search">
    <input tabindex="2" type="submit" name="submit" value="Submit" id="submit">


  </form>
  <h2>Inspirational Quotes</h2>
  <blockquote>
    <p>&ldquo;There's no Theory of Evolution, just a list of creatures I've allowed to live.&rdquo;<br>
    - Chuck Norris</p>
  </blockquote>
  <blockquote>
    <p>&ldquo;Wise men say forgiveness is divine, but never pay full price for late pizza.&rdquo;<br>
    - TMNT</p>
  </blockquote>
  <footer>&copy; 2018 Camper Cat</footer>
</body>
<style>
  body {
    height: 100%;
    margin: 0 !important;
    padding: 8px;
  }
  .overlay {
    margin: -8px;
    position: absolute;
    width: 100%;
    height: 100%;
  }
</style>
```
