---
id: 587d774c367417b2b2512a9c
title: Fügen Sie eine Textalternative für Bilder zur Barrierefreiheit bei Sehbeeinträchtigungen hinzu
challengeType: 0
videoUrl: 'https://scrimba.com/c/cPp7VfD'
forumTopicId: 16628
dashedName: add-a-text-alternative-to-images-for-visually-impaired-accessibility
---

# --description--

Sie haben wahrscheinlich schon ein `alt`-Attribut in einem `img`-Tag in anderen Aufgaben gesehen. `alt`-Text beschreibt den Inhalt des Bildes und bietet eine Textalternative dazu. Ein `alt`-Attribut hilft in Fällen, in denen das Bild nicht geladen wird oder von einem Benutzer nicht gesehen werden kann. Suchmaschinen verwenden es auch, um zu verstehen, was ein Bild enthält, damit es in Suchergebnissen angezeigt wird. Hier ist ein Beispiel:

```html
<img src="importantLogo.jpeg" alt="Company logo">
```

Menschen mit Sehbeeinträchtigungen sind auf Screenreader angewiesen, die Webinhalte in eine Audio-Schnittstelle umwandeln. Sie erhalten keine Informationen, wenn diese nur visuell präsentiert werden. Für Bilder können Screenreader auf das `alt`-Attribut zugreifen und dessen Inhalt vorlesen, um wichtige Informationen zu vermitteln.

Guter `alt`-Text bietet dem Leser eine kurze Beschreibung des Bildes. Sie sollten immer ein `alt`-Attribut an Ihrem Bild hinzufügen. Laut HTML5-Spezifikation gilt dies inzwischen als verpflichtend.

# --instructions--

Camper Cat ist sowohl ein Coding-Ninja als auch ein echter Ninja, der eine Website erstellt, um sein Wissen zu teilen. Das Profilbild, das er verwenden möchte, zeigt seine Fähigkeiten und sollte von allen Seitenbesuchern geschätzt werden. Fügen Sie ein `alt`-Attribut im `img`-Tag hinzu, das erklärt, dass Camper Cat Karate macht. (Das Bild-`src` verweist nicht auf eine tatsächliche Datei, daher sollten Sie den `alt`-Text in der Anzeige sehen.)

# --hints--

Ihr `img`-Tag sollte ein `alt`-Attribut haben, und dieses darf nicht leer sein.

```js
assert.isNotEmpty(document.querySelector('img')?.getAttribute('alt'));
```

# --seed--

## --seed-contents--

```html
<img src="doingKarateWow.jpeg">
```

# --solutions--

```html
<img src="doingKarateWow.jpeg" alt="Someone doing karate">
```
