---
id: 587d778f367417b2b2512aae
title: Links Bedeutung verleihen durch beschreibenden Linktext
challengeType: 0
videoUrl: 'https://scrimba.com/c/c437DcV'
forumTopicId: 301013
dashedName: give-links-meaning-by-using-descriptive-link-text
---

# --description--

Screenreader-Nutzende haben verschiedene Optionen, welche Art von Inhalten ihr Gerät vorliest. Diese Optionen umfassen das Überspringen zu (oder über) Landmark-Elemente, das Springen zum Hauptinhalt oder das Abrufen einer Seitenübersicht anhand der Überschriften. Eine weitere Möglichkeit ist, nur die auf einer Seite verfügbaren Links anzuhören.

Screenreader tun dies, indem sie den Linktext vorlesen, also den Text zwischen den Anker- (`a`) Tags. Eine Liste mit Links wie „hier klicken“ oder „mehr lesen“ ist nicht hilfreich. Stattdessen verwenden Sie kurzen, aber beschreibenden Text innerhalb der `a` Tags, um diesen Nutzenden mehr Bedeutung zu vermitteln.

# --instructions--

Der Linktext, den Camper Cat verwendet, ist ohne den umgebenden Kontext nicht sehr aussagekräftig. Verschieben Sie die Anker- (`a`) Tags so, dass sie den Text `information about batteries` umschließen, anstatt `Click here`.

# --hints--

Ihr Code sollte die Anker-`a` Tags von den Worten `Click here` wegbewegen und stattdessen um die Worte `information about batteries` legen.

```js
assert.match(document.querySelector('a')?.textContent, /^(information about batteries)$/g);
```

Das `a` Element sollte ein `href` Attribut mit dem Wert eines leeren Strings `""` haben.

```js
assert.isEmpty(document.querySelector('a')?.getAttribute('href'));
```

Das `a` Element sollte ein schließendes Tag haben.

```js
assert.isTrue(code.match(/<\/a>/g)?.length === code.match(/<a href=(''|"")>/g)?.length);
```

# --seed--

## --seed-contents--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. <a href="">Click here</a> for information about batteries</p>
  </article>
</body>
```

# --solutions--

```html
<body>
  <header>
    <h1>Deep Thoughts with Master Camper Cat</h1>
  </header>
  <article>
    <h2>Defeating your Foe: the Red Dot is Ours!</h2>
    <p>Felines the world over have been waging war on the most persistent of foes. This red nemesis combines both cunning stealth and lightning speed. But chin up, fellow fighters, our time for victory may soon be near. Click here for <a href="">information about batteries</a></p>
  </article>
</body>
```
