---
id: 61537485c4f2a624f18d7794
title: 步骤 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Begin with your standard HTML boilerplate. Add a `DOCTYPE` declaration, an `html` element, a `head` element, and a `body` element.

将值为 `en` 的 `lang` 属性添加到 `<html>` 开始标签里。

# --hints--

代码应包含 `DOCTYPE` 声明。

```js
assert(code.match(/<!DOCTYPE/gi));
```

应该在 `DOCTYPE` 声明之后包含一个空格。

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

应该定义文档类型为 `html`。

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

应该在 `DOCTYPE` 声明的类型后面加上一个 `>`。

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

`html` 元素应该有一个开始标签。

```js
assert(code.match(/<html\s*>|<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

`html` 元素应该有一个结束标签。

```js
assert(code.match(/<\/html\s*>/));
```

`<html>` 开始标签应该有一个值为 `en` 的 `lang` 属性。

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/));
```

您的 `DOCTYPE` 声明应位于最前面（HTML的开头）。

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

您应有一个 `head` 开始标签。

```js
assert(code.match(/<head\s*>/i));
```

您应有一个 `head` 结束标签。

```js
assert(code.match(/<\/head\s*>/i));
```

您应有一个 `body` 开始标签。

```js
assert(code.match(/<body\s*>/i));
```

您应有一个 `body` 结束标签。

```js
assert(code.match(/<\/body\s*>/i));
```

`head` 和 `body` 元素应该是同级元素。

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

`head` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

`body` 元素应该在 `html` 元素内。

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
