---
id: 63b60aaaa65f8922bfce6b7e
title: Крок 9
challengeType: 0
dashedName: step-9
---

# --description--

У вашому елементі `span` створіть елемент `label` для `entry-dropdown` і додайте йому текст `Add food or exercise:`. Потім створіть елемент `select` з атрибутом `id`, встановленим у `entry-dropdown`, і атрибутом `name`, встановленим у `options`. Нижче додайте елемент `button` з атрибутом `id`, встановленим у `add-entry`, і текстом `Add Entry`.

Додайте до вашого елемента `button` атрибут `type`, встановлений у `button`, щоб запобігти автоматичній відправці форми.

# --hints--

Ви повинні додати елемент `label` до вашого елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > label'));
```

Ваш новий елемент `label` має мати атрибут `for`, встановлений у `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > label')?.getAttribute('for'), 'entry-dropdown');
```

Ваш новий елемент `label` має містити текст `Add food or exercise:`.

```js
assert.equal(document.querySelector('.controls > span > label')?.innerText, 'Add food or exercise:');
```

Ви повинні додати елемент `select` до вашого елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > select'));
```

Ваш елемент `select` має йти після елемента `label`.

```js
assert(document.querySelector('.controls > span > select')?.previousElementSibling?.tagName === 'LABEL');
```

Ваш новий елемент `select` має мати атрибут `id`, встановлений у `entry-dropdown`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('id'), 'entry-dropdown');
```

Ваш новий елемент `select` має мати атрибут `name`, встановлений у `options`.

```js
assert.equal(document.querySelector('.controls > span > select')?.getAttribute('name'), 'options');
```

Ваш елемент `select` має бути порожнім і мати закриваючий тег `</select>`.

```js
assert.equal(document.querySelector('.controls > span > select')?.innerHTML?.trim(), '');
```

Ви повинні додати елемент `button` до вашого елемента `span`.

```js
assert.exists(document.querySelector('.controls > span > button'));
```

Ваш елемент `button` має йти після елемента `select`.

```js
assert(document.querySelector('.controls > span > button')?.previousElementSibling?.tagName === 'SELECT');
```

Ваш новий елемент `button` має мати атрибут `id`, встановлений у `add-entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.getAttribute('id'), 'add-entry');
```

Ваш новий елемент `button` має містити текст `Add Entry`.

```js
assert.equal(document.querySelector('.controls > span > button')?.innerText, 'Add Entry');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <link rel="stylesheet" href="styles.css" />
    <title>Calorie Counter</title>
  </head>
  <body>
    <main>
      <h1>Calorie Counter</h1>
      <div class="container">
        <form id="calorie-counter">
          <label for="budget">Budget</label>
          <input
            type="number"
            min="0"
            id="budget"
            placeholder="Daily calorie budget"
            required
          />
          <fieldset id="breakfast">
            <legend>Breakfast</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="lunch">
            <legend>Lunch</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="dinner">
            <legend>Dinner</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="snacks">
            <legend>Snacks</legend>
            <div class="input-container"></div>
          </fieldset>
          <fieldset id="exercise">
            <legend>Exercise</legend>
            <div class="input-container"></div>
          </fieldset>
--fcc-editable-region--
          <div class="controls">
            <span>

            </span>
          </div>
--fcc-editable-region--
        </form>
      </div>
    </main>
  </body>
</html>
```

```css
:root {
  --light-grey: #f5f6f7;
  --dark-blue: #0a0a23;
  --fcc-blue: #1b1b32;
  --light-yellow: #fecc4c;
  --dark-yellow: #feac32;
  --light-pink: #ffadad;
  --dark-red: #850000;
  --light-green: #acd157;
}

body {
  font-family: "Lato", Helvetica, Arial, sans-serif;
  font-size: 18px;
  background-color: var(--fcc-blue);
  color: var(--light-grey);
}

h1 {
  text-align: center;
}

.container {
  width: 90%;
  max-width: 680px;
}

h1,
.container,
.output {
  margin: 20px auto;
}

label,
legend {
  font-weight: bold;
}

.input-container {
  display: flex;
  flex-direction: column;
}

button {
  cursor: pointer;
  text-decoration: none;
  background-color: var(--light-yellow);
  border: 2px solid var(--dark-yellow);
}

button,
input,
select {
  min-height: 24px;
  color: var(--dark-blue);
}

fieldset,
label,
button,
input,
select {
  margin-bottom: 10px;
}

.output {
  border: 2px solid var(--light-grey);
  padding: 10px;
  text-align: center;
}

.hide {
  display: none;
}

.output span {
  font-weight: bold;
  font-size: 1.2em;
}

.surplus {
  color: var(--light-pink);
}

.deficit {
  color: var(--light-green);
}
```

```js

```
