---
id: 669e81368e52b3a5c35a2dc5
title: Створіть сторінку з відеокомпіляцією
challengeType: 25
dashedName: build-a-video-compilation-page
demoType: onClick
---

# --description--

**Мета:** Виконайте наведені нижче історії користувача та пройдіть усі тести, щоб завершити лабораторну роботу.

**Історії користувача:**

1. Ви повинні мати елемент `main` як єдиний дочірній елемент `body`.
1. Ви повинні мати елемент `h1` з темою вашої сторінки.
1. Під вашим елементом `h1` має бути абзац, який вводить тему вашої сторінки.
1. Під вашим першим абзацом має бути три елементи `section`.
1. Кожен розділ має містити елемент `h2`, абзац і елемент `iframe` у такому порядку.
1. Три елементи `iframe` повинні мати атрибут `src`, встановлений на дійсне відео.
1. Кожен елемент `iframe` також повинен мати атрибут `title` для опису вставленого вмісту, а також атрибути `height` і `width`, щоб встановити елемент у відповідний розмір.

# --hints--

Ви повинні мати елемент `main` всередині вашого елемента `body`.

```js
assert.equal(document.querySelector('body :first-child').tagName, 'MAIN');
```

Ваш елемент `main` має бути єдиним дочірнім елементом `body`.

```js
assert.equal(document.querySelector('body').children.length, 1);
```

Ви повинні мати елемент `h1` з темою вашої сторінки всередині елемента `main`.

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

Під вашим елементом `h1` має бути абзац, який вводить тему вашої сторінки.

```js
const p = document.querySelectorAll('p')[0];
assert.equal(p?.previousElementSibling?.tagName, 'H1');
assert.isAbove(p?.innerText.length, 0);
```

Під вашим першим елементом `section` має бути три елементи `p`.

```js
assert.equal(document.querySelectorAll('section')?.length, 3);
assert.equal(document.querySelectorAll('section')[0]?.previousElementSibling?.tagName, 'P')
```

Кожен елемент `section` має починатися з елемента `h2`, який слугує заголовком для цього розділу.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[0]?.tagName, 'H2');
    assert.isAbove(section.children[0]?.innerText.length, 0);
}
```

Кожен елемент `section` має містити елемент `p` як другий дочірній елемент, який вводить відеоконтент.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[1]?.tagName, 'P');
    assert.isAbove(section.children[1]?.innerText.length, 0);
}
```

Кожен елемент `section` має містити елемент `iframe` як третій дочірній елемент.

```js
const sections = document.querySelectorAll('section');
assert.equal(sections?.length, 3)
for (let section of sections) {
    assert.equal(section.children[2]?.tagName, 'IFRAME');
}
```

Ваш перший елемент `iframe` має мати атрибут `src`, встановлений на дійсне відео.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.src.length, 0)
```

Ваш перший елемент `iframe` має мати атрибут `title` для опису вставленого вмісту.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.title.length, 0)
```

Ваш перший елемент `iframe` має мати атрибут `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.height.length, 0)
```

Ваш перший елемент `iframe` має мати атрибут `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[0]?.width.length, 0)
```

Ваш другий елемент `iframe` має мати атрибут `src`, встановлений на дійсне відео.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.src.length, 0)
```

Ваш другий елемент `iframe` має мати атрибут `title` для опису вставленого вмісту.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.title.length, 0)
```

Ваш другий елемент `iframe` має мати атрибут `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.height.length, 0)
```

Ваш другий елемент `iframe` має мати атрибут `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[1]?.width.length, 0)
```

Ваш третій елемент `iframe` має мати атрибут `src`, встановлений на дійсне відео.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.src.length, 0)
```

Ваш третій елемент `iframe` має мати атрибут `title` для опису вставленого вмісту.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.title.length, 0)
```

Ваш третій елемент `iframe` має мати атрибут `height`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.height.length, 0)
```

Ваш третій елемент `iframe` має мати атрибут `width`.

```js
assert.isAbove(document.querySelectorAll('iframe')[2]?.width.length, 0)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="utf-8">
    <title>Video Compilation Page</title>
</head>

<body>
    <main>
        <h1>Front-End Web Development</h1>
        <p>Front-End Web Development involves designing and building user interfaces of websites using HTML, CSS, and
            JavaScript. It focuses on creating visually appealing, interactive, and user-friendly websites. Front-end
            developers ensure accessibility, performance optimization, and
            adherence to web standards, blending creativity with technical expertise.</p>
        <section>
            <h2>HTML</h2>
            <p>HTML, or HyperText Markup Language, forms the backbone of web pages by defining their basic structure and
                layout. It uses a series of elements to display text, images, lists, tables, and other static content on
                a web page.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/GDGejH3SDNQ?si=KJYLgcz4kyyroYMB"
                title="html and coding introduction video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>CSS</h2>
            <p>CSS, or Cascading Style Sheets, is used to control the visual presentation of HTML elements on a web
                page. It allows for customized styling, including colors, fonts, layouts, and spacing, making websites
                visually appealing. CSS also supports responsive design, ensuring web pages look good on
                all devices.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/OXGznpKZ_sA?si=s3KDSZvrhU_PU9XL"
                title="css tutorial video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
        <section>
            <h2>JavaScript</h2>
            <p>JavaScript is a programming language that adds interactivity to static web pages, enabling dynamic
                content updates, form validations, animations, and responsive behaviors based on user interactions and
                events.</p>
            <iframe width="560" height="315" src="https://www.youtube.com/embed/jS4aFq5-91M?si=zKQhHEYwU4tnMmVm"
                title="javascript programming video" frameborder="0"
                allow="accelerometer; autoplay; clipboard-write; encrypted-media; gyroscope; picture-in-picture; web-share"
                referrerpolicy="strict-origin-when-cross-origin" allowfullscreen></iframe>
        </section>
    </main>
</body>

</html>
```
