---
id: 694648acde178bb8202d9517
title: Paso 10
challengeType: 20
dashedName: step-10
---

# --description--

La salida es `True`, confirmando que `score` es un `float`.

Completa la boleta imprimiendo el valor de `score` junto con su tipo de dato usando una sola instrucción `print()`.

# --hints--

Debes imprimir `score` y `type(score)` usando un separador de coma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
--fcc-editable-region--

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
