---
id: 694648acde178bb8202d9512
title: Paso 4
challengeType: 20
dashedName: step-4
---

# --description--

La boleta de calificaciones también debe mostrar si el estudiante está actualmente inscrito. Esto puede representarse usando un valor booleano.

Los valores booleanos representan una condición de sí o no, y a menudo se usan para tomar decisiones en el código. Solo existen dos valores booleanos: `True` y `False`.

Declara una variable llamada `is_student` y asígnale el valor `True`.

# --hints--

Deberías tener una variable `is_student`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("is_student")`) })
```

La variable `is_student` debe almacenar el valor `True`. No rodees el valor con comillas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("is_student").is_equivalent("is_student = True")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

--fcc-editable-region--

--fcc-editable-region--
```
