---
id: 694648acde178bb8202d9513
title: Paso 6
challengeType: 20
dashedName: step-6
---

# --description--

El nombre del estudiante debe seguir el mismo formato que los otros detalles.

Elimina las salidas anteriores de la variable `name`. Luego, imprime `name` y `type(name)` juntos en una línea separados por una coma como en el paso anterior.

# --hints--

No deberías tener `print(name)` en tu código.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(name)")`)})
```

No deberías tener `print(type(name))` en tu código.

```js
({ test: () => runPython(`assert not _Node(_code).has_call("print(type(name))")`)})
```

Debes imprimir `name` y `type(name)` en la misma llamada a `print()` usando una coma como separador.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name, type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--
print(name)
print(type(name))
--fcc-editable-region--

is_student = True
print(is_student, type(is_student))
```
