---
id: 694648acde178bb8202d9514
title: Paso 7
challengeType: 20
dashedName: step-7
---

# --description--

Ahora necesitas agregar la edad del estudiante a la boleta de calificaciones. Para eso usarás un entero, uno de los tipos de datos numéricos en Python.

Declara una variable llamada `age` y asígnale el valor entero `20`.

Luego, imprime el valor y el tipo de dato de `age` juntos separados por una coma. Revisa la salida en el terminal que muestra el valor de `age` y su tipo como `int` (entero).

# --hints--

Deberías tener una variable `age`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("age")`) })
```

La variable `age` debe almacenar el valor `20`. No rodees el valor con comillas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("age").is_equivalent("age = 20")`) })
```

Debes imprimir `age` y `type(age)` usando un separador de coma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(age, type(age))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

--fcc-editable-region--

--fcc-editable-region--

```
