---
id: 69461e4f98594012939bc1c6
title: Paso 1
challengeType: 20
dashedName: step-1
---

# --description--

En este taller, practicarás los tipos de datos de Python construyendo una impresora de boletas simple.

Como aprendiste en lecciones anteriores, las variables se asignan de esta manera:

```py
greeting = 'Hello World'
```

En el ejemplo anterior, la variable `greeting` tiene el valor de una cadena, que es una secuencia de caracteres rodeada por comillas.

Crea una variable llamada `name` y asígnale la cadena `Alice`. Recuerda rodear el valor con comillas simples o dobles como se muestra en el ejemplo.

**Importante:** Python usa la indentación (los espacios al inicio de una línea) para organizar el código. Para este taller, asegúrate de que no haya espacios adicionales al comienzo de cada línea de código. Agregar espacios extra causará un `IndentationError` y evitará que tu código se ejecute.

# --hints--

Deberías tener una variable `name`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable('name')`) })
```

La variable `name` debe almacenar el valor `'Alice'`.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("name").is_equivalent("name = 'Alice'")`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
