---
id: 694648acde178bb8202d9516
title: Paso 9
challengeType: 20
dashedName: step-9
---

# --description--

La salida es `False`, lo que muestra que `score` no es un `int`.

Otro tipo común de número en Python es `float`, que representa un número con decimales. Reemplaza `int` con `float` en la llamada existente a `isinstance()` para confirmar esto.

Finalmente, imprime `score` y su tipo de dato para completar la boleta de calificaciones.

# --hints--

Debes reemplazar `int` con `float` en la llamada existente `isinstance(score, int)`.

```js
({ test: () => runPython(`
  assert not _Node(_code).has_call("print(isinstance(score, int))")
  assert _Node(_code).has_call("print(isinstance(score, float))")
`) })
```

Debes imprimir `score` y `type(score)` usando un separador de coma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(score, type(score))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
--fcc-editable-region--
print(isinstance(score, int))

--fcc-editable-region--

```

# --solutions--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

score = 80.5
print(isinstance(score, float))
print(score, type(score))
```
