---
id: 694648abde178bb8202d9511
title: Paso 3
challengeType: 20
dashedName: step-3
---

# --description--

Ahora deberías ver el nombre del estudiante impreso en el terminal.

Python proporciona una función llamada `type()` que puedes usar para verificar el tipo de un valor.

```py
platform = 'freeCodeCamp'
print(type(platform)) # Output: <class 'str'>
```

En el ejemplo anterior, la salida `<class 'str'>` significa que la variable pasada a la función `type()` es una cadena.

Usa la función `type()` con `name` como argumento e imprime la salida como en el ejemplo. Revisa la salida en el terminal que muestra que `name` es del tipo `str` (cadena).

# --hints--

Debes imprimir el resultado de llamar a `type(name)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(type(name))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
--fcc-editable-region--

--fcc-editable-region--
```
