---
id: 694648acde178bb8202d9515
title: Paso 8
challengeType: 20
dashedName: step-8
---

# --description--

Ahora, añade la puntuación del estudiante.

Declara una variable llamada `score` y asígnale el valor `80.5`.

Aunque tanto `age` como `score` son números, pueden no ser del mismo tipo. Python proporciona una función llamada `isinstance()` para verificar esto.

```py
x = 10
print(isinstance(x, int)) # Output: True
```

Usa `isinstance()` para verificar si `score` es un `int`, y muestra el resultado en el terminal como se muestra en el ejemplo anterior.

# --hints--

Debes tener una variable `score`.

```js
({ test: () => runPython(`assert _Node(_code).has_variable("score")`) })
```

La variable `score` debe almacenar el valor `80.5`. No rodees el valor con comillas.

```js
({ test: () => runPython(`assert _Node(_code).find_variable("score").is_equivalent("score = 80.5")`) })
```

Debes imprimir el resultado de `isinstance(score, int)`.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(isinstance(score, int))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name, type(name))

is_student = True
print(is_student, type(is_student))

age = 20
print(age, type(age))

--fcc-editable-region--

--fcc-editable-region--

```
