---
id: 6947a3e516bb8a4a681274a7
title: Paso 5
challengeType: 20
dashedName: step-5
---

# --description--

La función `print()` puede mostrar más de un valor a la vez. Separa los valores con una coma (`,`) para imprimirlos en la misma línea.

```py
subject = 'Python'
print(subject, type(subject)) # Output: Python <class 'str'>
```

Imprime tanto `is_student` como `type(is_student)` en la misma línea usando una coma `,` como se muestra en el ejemplo. Luego, verifica la salida en la terminal que muestra el valor de `is_student` y su tipo como `bool` (booleano).

# --hints--

Debes imprimir `is_student` y `type(is_student)` usando un separador de coma.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(is_student, type(is_student))")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
print(name)
print(type(name))

is_student = True
--fcc-editable-region--

--fcc-editable-region--
```
