---
id: 6947a0c2b6de035dd8d3eea1
title: Paso 2
challengeType: 20
dashedName: step-2
---

# --description--

Puedes imprimir el valor de una variable usando la función `print()`.

```py
greeting = 'Hello World'
print(greeting) # Output: Hello World
```

Aprenderás más sobre funciones en las próximas lecciones. Por ahora, sabe que una función es un bloque reutilizable de código que puede ser llamado, o invocado, para ejecutar su código, y se le pueden pasar argumentos.

En el ejemplo anterior, `print(greeting)` es una llamada a función, y `greeting` es el argumento de la función.

Consulta el ejemplo e imprime la variable `name`. Revisa la salida en el terminal.

# --hints--

Debes llamar a `print()` con `name` como argumento.

```js
({ test: () => runPython(`assert _Node(_code).has_call("print(name)")`) })
```

# --seed--

## --seed-contents--

```py
name = 'Alice'
--fcc-editable-region--

--fcc-editable-region--
```
