---
id: 68b94f23f7b3ee14078c89c5
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

En este taller, usarás el elemento `iframe` para mostrar un video. La estructura básica de HTML ya ha sido preparada para ti.

Comienza creando un elemento `h1` con el texto `iframe Video Display`.

# --hints--

Debería tener un elemento `h1`.

```js
assert.exists(document.querySelector('h1'))
```

Tu elemento `h1` debe tener el texto `iframe Video Display`.

```js
assert.equal(document.querySelector('h1')?.textContent.trim(), 'iframe Video Display')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
