---
id: 68ba9c4f1688914d72876095
title: Paso 4
challengeType: 0
dashedName: step-4
---

# --description--

El elemento `iframe` también toma un atributo `src` con un valor que indica la URL o la ruta del recurso a mostrar.

Agrega un atributo `src` de `https://www.youtube.com/embed/I0_951_MPE0` a tu elemento `iframe`.

En este punto, deberías ver el video mostrando en la página, pero hay algunos atributos más que necesitas agregar.

# --hints--

Tu elemento `iframe` debe tener un atributo `src` establecido en `https://www.youtube.com/embed/I0_951_MPE0`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('src'), 'https://www.youtube.com/embed/I0_951_MPE0')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    <iframe
      width="560"
      height="315"

    >

    </iframe>
--fcc-editable-region--
  </body>
</html>
```
