---
id: 68ba9c4f1688914d72876096
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Uno de los atributos es `allow`. Es como una lista de permisos que le dice al navegador qué funciones puede usar el `iframe`.

Aquí tienes un elemento `iframe` con el atributo `allow`:

```html
<iframe
  allow="accelerometer autoplay clipboard-write encrypted-media gyroscope picture-in-picture web-share"
></iframe>
```

Agrega el atributo `allow` con el valor `accelerometer`, `autoplay` y `clipboard-write`.

`accelerometer` permite que el `iframe` use sensores de movimiento para que pueda detectar cosas como la inclinación y rotación del dispositivo. `autoplay` permite que el video comience a reproducirse automáticamente, y `clipboard-write` permite que el iframe escriba datos en el portapapeles del usuario.

# --hints--

Tu elemento `iframe` debe tener un atributo `allow` configurado en `accelerometer`.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /accelerometer/)
```

El atributo `allow` de tu elemento `iframe` debe tener `autoplay` como uno de sus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /autoplay/)
```

El atributo `allow` de tu elemento `iframe` debe tener `clipboard-write` como uno de sus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /clipboard-write/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
