---
id: 68ba9c4f1688914d72876094
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

En la primera lección sobre el elemento `iframe`, aprendiste que es un elemento reemplazado al igual que `img`. Eso significa que también puede tomar las propiedades `width` y `height` para determinar qué tan alto y ancho debe ser.

Dale a tu elemento `iframe` un `width` de `560` y un `height` de `315`.

# --hints--

Tu elemento `iframe` debe tener un atributo `width` establecido en `560`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('width'), '560')
```

Tu elemento `iframe` debe tener un atributo `height` establecido en `315`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('height'), '315')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>

    <iframe
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
