---
id: 68ba9c4f1688914d72876098
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

El siguiente atributo que agregarás es `referrerpolicy`. Es la regla que determina cuánta información compartes cuando tu página se conecta a otra página.

Agrega el atributo `referrerpolicy` y configúralo en `strict-origin-when-cross-origin`. Esto comparte la dirección completa en el mismo sitio, solo el nombre del sitio en otros sitios, y nada en sitios inseguros.

# --hints--

Tu elemento `iframe` debe tener un atributo `referrerpolicy` configurado en `strict-origin-when-cross-origin`.

```js
const iframeEl = document.querySelector('iframe')
assert.equal(iframeEl?.getAttribute('referrerpolicy'), 'strict-origin-when-cross-origin')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
      allow="accelerometer autoplay clipboard-write encrypted-media gyroscope web-share"
--fcc-editable-region--
      
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
