---
id: 68ba9c4f1688914d72876093
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Ahora, crea un elemento `iframe`. No pongas nada dentro todavía.

# --hints--

Debes tener un elemento `iframe`.

```js
assert.exists(document.querySelector('iframe'))
```

Tu elemento `iframe` debe estar vacío.

```js
assert.equal(document.querySelector('iframe')?.innerHTML.trim(), '')
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
--fcc-editable-region--
    
--fcc-editable-region--
  </body>
</html>
```
