---
id: 68ba9c4f1688914d72876097
title: Paso 6
challengeType: 0
dashedName: step-6
---

# --description--

Agrega `encrypted-media`, `gyroscope` y `web-share` a los valores existentes en el atributo `allow`.

Estos tres permitirán el uso de extensiones de medios cifrados para proteger el video, otorgar acceso a los sensores de movimiento y orientación del dispositivo, y permitir compartir el contenido del iframe a través de los diálogos nativos de compartir del dispositivo. 

# --hints--

El atributo `allow` de tu elemento `iframe` debe tener `encrypted-media` como uno de sus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /encrypted-media/)
```

El atributo `allow` de tu elemento `iframe` debe tener `gyroscope` como uno de sus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /gyroscope/)
```

El atributo `allow` de tu elemento `iframe` debe tener `web-share` como uno de sus valores.

```js
const iframeEl = document.querySelector('iframe')
const iframeElAllowAttr = iframeEl?.getAttribute('allow')
assert.match(iframeElAllowAttr, /web-share/)
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Display Videos in an iframe</title>
  </head>
  <body>
    <h1>iframe Video Display</h1>
    <iframe
      width="560"
      height="315"
      src="https://www.youtube.com/embed/I0_951_MPE0"
--fcc-editable-region--
      allow="accelerometer autoplay clipboard-write"
--fcc-editable-region--
    >

    </iframe>
  </body>
</html>
```
