---
id: 6909ed80e1c195f49564edac
title: Hatua 1
challengeType: 1
dashedName: step-1
---

# --description--

Katika warsha hii, utafundishwa jinsi ya kufanya kazi na sentensi za masharti na vionyeshi vya kulinganisha.

Anza kwa kutangaza kigezo kinachoitwa `truthyOrFalsy` na kisha weka thamani ya kweli au siyo kweli `true`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Unapaswa kuweka thamani `true` kwa kigezo chako cha `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
