---
id: 690a51806fb743bc896c3af7
title: Hatua 2
challengeType: 1
dashedName: step-2
---

# --description--

Kama ulivyojifunza katika funzo lililopita, thamani za kweli au siyo kweli ni aina ya data yenye thamani za `true` na `false` pekee.

Thamani ya kweli ni thamani inayotathminiwa kuwa `true` katika muktadha wa thamani ya kweli au siyo kweli. Mifano ya thamani za kweli ni:

- Mfuatano wa herufi usio tupu
- Nambari zisizo sifuri
- Thamani ya kweli au siyo kweli `true`

Na thamani nyingine utakazokutana nazo baadaye.

Unaweza kutumia kitendakazi cha `Boolean()` kukagua kama thamani ni ya kweli. Kwa mfano, `Boolean("Hello World!")` itarudisha `true` kwa sababu `"Hello World!"` ni thamani ya kweli.

Sasa, fanya wito wa `Boolean(truthyOrFalsy)` na andika matokeo kwenye konsoli.

# --hints--

Unapaswa kuandika `Boolean(truthyOrFalsy)` kwenye konsoli.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
