---
id: 690a0df27ee9ee36d935d530
title: Hatua 8
challengeType: 1
dashedName: step-8
---

# --description--

Tangaza kigezo kingine kinachoitwa `isTimmyAGamer` na weka thamani yake kuwa thamani ya kweli au siyo kweli `false`.

# --hints--

Unapaswa kutangaza kigezo kinachoitwa `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Unapaswa kuweka thamani ya kweli au siyo kweli `false` kwa kigezo chako cha `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
