---
id: 690a06833d3e83b5226ade78
title: Hatua ya 7
challengeType: 1
dashedName: step-7
---

# --description--

Usemi wa masharti hutumika kuendesha msimbo kulingana na hali fulani. Unaundwa na neno kuu la `if`, linalofuatiwa na hali iliyozungushwa na mabano ya kawaida, na mwili uliowekwa ndani ya mabano ya mviringo. Hapa kuna mfano wa usemi wa `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Wakati `condition` ni thamani ya kweli (mfuatano wa herufi usio tupu, nambari isiyo sifuri, `true`, n.k.) msimbo uliomo ndani ya mwili wa usemi wa `if` utaendeshwa.

Chini ya kigezo chako, tengeneza usemi wa `if` unaotumia `hasDeveloperJob` kama hali yake. Ndani ya mwili wa usemi wako wa `if`, andika `"Timmy is employed as a developer."` kwenye konsoli.

# --hints--

Unapaswa kuwa na usemi wa `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Usemi wako wa `if` unapaswa kutumia `hasDeveloperJob` kama hali yake.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Unapaswa kuandika `"Timmy is employed as a developer."` kwenye konsoli ndani ya usemi wako wa `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
