---
id: 690a0f7a5c22c15222a6aa5c
title: Hatua 9
challengeType: 1
dashedName: step-9
---

# --description--

Sasa tengeneza tamko jingine la `if` ambalo linatumia `isTimmyAGamer` kama sharti lake. Ndani ya mwili wa tamko lako la `if`, andika `"Timmy loves to play World of Warcraft."` kwenye konsoli.

Kumbuka kuwa hutakuwaona chochote kikiwa kimeandikwa kwenye konsoli kwa sababu sharti la tamko lako jipya la `if` ni false.

# --hints--

Unapaswa kuwa na tamko mbili za `if` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Tamko lako la pili la `if` linapaswa kutumia `isTimmyAGamer` kama sharti lake.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Unapaswa kuandika `"Timmy loves to play World of Warcraft."` kwenye konsoli ndani ya tamko lako la pili la `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
