---
id: 690a562fd611e8155599ce98
title: Hatua 3
challengeType: 1
dashedName: step-3
---

# --description--

Badilisha mgawo wa thamani wa `truthyOrFalsy` kwa kuweka thamani yake kuwa `"freeCodeCamp"`. Bado utaona `true` ikirekodiwa kwenye konsoli kwa sababu mfuatano wa herufi usio tupu ni thamani ya kweli.

# --hints--

Kigezo chako cha `truthyOrFalsy` kinapaswa kuwa na thamani `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
