---
id: 690a3ad44874a928a2fe3077
title: Hatua 11
challengeType: 1
dashedName: step-11
---

# --description--

Viwango vya kulinganisha vinakuwezesha kulinganisha thamani na kuandika kauli za masharti tata zaidi.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Kwa mfano, msimbo ulio juu unaandika `"Happy Birthday!"` kwenye konsoli tu wakati thamani za vigezo `today` na `birthday` ni sawa.

Tengeneza kauli ya tatu ya `if`. Kwa sharti lake, tumia kiendeshaji `>=` kuangalia kama `timmyAge` ni kubwa au sawa na `16`.

Ndani ya mwili wa kauli yako mpya ya `if`, andika `"Timmy is old enough to drive."` kwenye konsoli.

# --hints--

Unapaswa kuwa na kauli ya tatu ya `if` katika msimbo wako.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Kauli yako ya tatu ya `if` inapaswa kutumia `timmyAge >= 16` kama sharti lake.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

Unapaswa kuandika `"Timmy is old enough to drive."` kwenye konsoli ndani ya kauli yako ya tatu ya `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
