---
id: 690a57fef66aac39dffc61a2
title: Hatua 5
challengeType: 1
dashedName: step-5
---

# --description--

Sasa unajua thamani za truthy na falsy ni zipi, ondoa tamko lako la variable pamoja na tamko la `console.log` kutoka kwenye msimbo wako.

# --hints--

Haupaswi kuwa na variable ya `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

Haupaswi kuwa na `console.log(Boolean(truthyOrFalsy))` katika msimbo wako.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
