---
id: 690a110a76d0496d9e7c28eb
title: Hatua 10
challengeType: 1
dashedName: step-10
---

# --description--

Tangaza kigezo kingine kinachoitwa `timmyAge` na weka thamani yake kuwa `18`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `timmyAge`.

```js
assert.exists(timmyAge);
```

Unapaswa kuweka thamani ya `18` katika kigezo chako cha `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
