---
id: 690a58cc4f828a49eab1b53a
title: Hatua 6
challengeType: 1
dashedName: step-6
---

# --description--

Tangaza kigezo kinachoitwa `hasDeveloperJob` na weka thamani yake kuwa thamani ya kweli au siyo kweli `true`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Unapaswa kuweka thamani `true` kwa kigezo chako cha `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
