---
id: 690a57095c4afb27bbe751b3
title: Hatua 4
challengeType: 1
dashedName: step-4
---

# --description--

Thamani za uongo ni thamani zinazotathminiwa kuwa `false` katika muktadha wa thamani ya kweli au siyo kweli. Mifano ya thamani za uongo ni:

- `""` (mfuatano wa herufi tupu)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Utaziona zote unapoendelea na kozi hii.

Kwa sasa, badilisha mgawo wa thamani wa `truthyOrFalsy` kwa kuweka thamani yake kuwa mfuatano wa herufi tupu. Baada ya hapo, utaona `false` ikichapishwa kwenye konsoli.

# --hints--

Unapaswa kuweka thamani ya mfuatano wa herufi tupu kwa `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
