---
id: 690a3ce03013be44d116136e
title: Hatua 12
challengeType: 1
dashedName: step-12
---

# --description--

Usemi wa masharti unaweza kuwa na kifungu cha `else`, ambacho hufanya msimbo wakati hali ya `if` ni ya uongo. Hapa kuna mfano wa usemi wa `if...else`:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

Ongeza kifungu cha `else` kwenye usemi wa `if` uliopo. Ndani ya mwili wa kifungu chako cha `else`, andika `"Timmy is not old enough to drive."` kwenye konsoli.

# --hints--

Usemi wako wa tatu wa `if` unapaswa kuwa na kifungu cha `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s*else\s*\{/s);
```

Unapaswa kuandika `"Timmy is not old enough to drive."` kwenye konsoli ndani ya mwili wa kifungu chako cha `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s*else\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\2\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
