---
id: 6909ed80e1c195f49564edac
title: Крок 1
challengeType: 1
dashedName: step-1
---

# --description--

У цьому практичному занятті ви навчитеся працювати з умовними інструкціями та операторами порівняння.

Почніть із оголошення змінної з іменем `truthyOrFalsy` і призначте їй булеве значення `true`.

# --hints--

У вас має бути змінна з іменем `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

Ви повинні призначити `true` вашій змінній `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
