---
id: 690a57095c4afb27bbe751b3
title: Крок 4
challengeType: 1
dashedName: step-4
---

# --description--

Хибні значення — це значення, які оцінюються як `false` у булевому контексті. Прикладами хибних значень є:

- `""` (порожній рядок)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

Ви зустрінете всі ці значення, проходячи цей курс.

Наразі змініть присвоєння `truthyOrFalsy`, встановивши його значення порожнім рядком. Після цього ви побачите, що `false` виводиться в консоль.

# --hints--

Ви повинні призначити порожній рядок `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
