---
id: 690a110a76d0496d9e7c28eb
title: Крок 10
challengeType: 1
dashedName: step-10
---

# --description--

Оголосіть іншу змінну з іменем `timmyAge` і призначте їй значення `18`.

# --hints--

У вас має бути змінна з іменем `timmyAge`.

```js
assert.exists(timmyAge);
```

Ви повинні призначити `18` вашій змінній `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
