---
id: 690a3ad44874a928a2fe3077
title: Крок 11
challengeType: 1
dashedName: step-11
---

# --description--

Оператори порівняння дозволяють порівнювати значення та писати складніші умовні інструкції.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

Наприклад, код вище виводить у консоль `"Happy Birthday!"` лише тоді, коли значення змінних `today` і `birthday` рівні.

Створіть третю інструкцію `if`. Для її умови використайте оператор `>=`, щоб перевірити, чи `timmyAge` більше або дорівнює `16`.

У тілі нової інструкції `if` виведіть у консоль `"Timmy is old enough to drive."`.

# --hints--

У вашому коді має бути третя інструкція `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

Умова вашої третьої інструкції `if` має використовувати `timmyAge >= 16`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

Ви маєте вивести у консоль `"Timmy is old enough to drive."` всередині третьої інструкції `if`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
