---
id: 690a0df27ee9ee36d935d530
title: Крок 8
challengeType: 1
dashedName: step-8
---

# --description--

Оголосіть іншу змінну з іменем `isTimmyAGamer` і призначте їй булеве значення `false`.

# --hints--

Ви повинні оголосити змінну з іменем `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

Ви повинні призначити булеве значення `false` вашій змінній `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
