---
id: 690a562fd611e8155599ce98
title: Крок 3
challengeType: 1
dashedName: step-3
---

# --description--

Змініть присвоєння `truthyOrFalsy`, встановивши його значення рівним `"freeCodeCamp"`. Ви все одно побачите `true` у консолі, тому що непорожній рядок є правдивим значенням.

# --hints--

Ваша змінна `truthyOrFalsy` має мати значення `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
