---
id: 690a58cc4f828a49eab1b53a
title: Крок 6
challengeType: 1
dashedName: step-6
---

# --description--

Оголосіть змінну з ім’ям `hasDeveloperJob` і призначте їй булеве значення `true`.

# --hints--

У вас має бути змінна з ім’ям `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

Ви повинні призначити `true` змінній `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
