---
id: 690a57fef66aac39dffc61a2
title: Крок 5
challengeType: 1
dashedName: step-5
---

# --description--

Тепер, коли ви знаєте, що таке істинні та хибні значення, видаліть зі свого коду оголошення змінної та інструкцію `console.log`.

# --hints--

У вас не повинно бути змінної `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

У вашому коді не повинно бути `console.log(Boolean(truthyOrFalsy))`.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
