---
id: 690a06833d3e83b5226ade78
title: Крок 7
challengeType: 1
dashedName: step-7
---

# --description--

Умовна інструкція використовується для запуску коду залежно від умови. Вона складається з ключового слова `if`, за яким слідує умова в круглих дужках, а також тіла в фігурних дужках. Ось приклад інструкції `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

Коли `condition` є правдивим значенням (непорожній рядок, ненульове число, `true` тощо), виконується код у тілі інструкції `if`.

Під своєю змінною створіть інструкцію `if`, яка використовує `hasDeveloperJob` як умову. У тілі вашої інструкції `if` виведіть у консоль `"Timmy is employed as a developer."`.

# --hints--

У вас має бути інструкція `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

Ваша інструкція `if` має використовувати `hasDeveloperJob` як умову.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

Ви маєте вивести `"Timmy is employed as a developer."` у консоль у тілі інструкції `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
