---
id: 690a51806fb743bc896c3af7
title: Крок 2
challengeType: 1
dashedName: step-2
---

# --description--

Як ви дізналися в попередньому уроці, булевий тип — це тип даних, який має лише `true` та `false` значення.

Правдиве значення — це значення, яке в булевому контексті оцінюється як `true`. Прикладами правдивих значень є:

- Непорожні рядки
- Числа, що не дорівнюють нулю
- Булевий `true`

І інші значення, з якими ви познайомитеся пізніше.

Ви можете використовувати функцію `Boolean()`, щоб перевірити правдивість значення. Наприклад, `Boolean("Hello World!")` поверне `true`, тому що `"Hello World!"` є правдивим.

Тепер викличте `Boolean(truthyOrFalsy)` і виведіть результат у консоль.

# --hints--

Ви маєте вивести `Boolean(truthyOrFalsy)` у консоль.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
