---
id: 690a0f7a5c22c15222a6aa5c
title: Крок 9
challengeType: 1
dashedName: step-9
---

# --description--

Тепер створіть ще одну інструкцію `if`, яка використовує `isTimmyAGamer` як свою умову. Усередині тіла вашої інструкції `if` виведіть `"Timmy loves to play World of Warcraft."` у консоль.

Зверніть увагу, що ви не побачите нічого виведеного в консоль, оскільки умова вашої нової інструкції `if` є хибною.

# --hints--

У вашому коді має бути дві інструкції `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

Друга інструкція `if` має використовувати `isTimmyAGamer` як умову.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

Ви маєте вивести `"Timmy loves to play World of Warcraft."` у консоль всередині другої інструкції `if`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
