---
id: 68507cadbf36aa089a84ce2d
title: 實作 Luhn 演算法
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

Luhn 演算法，也稱為「modulus 10」或「mod 10」演算法，是一種簡單的校驗和公式，用來驗證各種識別號碼，例如信用卡號碼。使用 Luhn 演算法驗證號碼的步驟如下：

- 從右邊開始，並排除最右邊的數字（檢查碼），將每隔一個數字的值加倍。
- 如果將一個數字乘以 2 的結果大於 `9`，請將各位數字相加以取得單一數字。或者，你也可以從結果中減去 `9`。
- 將所有數字（包括檢查碼）相加。
- 如果所有數字的總和是 `10` 的倍數，則該數字有效；否則無效。

假設我們有數字 `453914881`。使用 Luhn 演算法驗證它的步驟如下：

 ```md 
 Account number      4   5   3   9   1   4   8   8   1 
 Double every other  4  10   3  18   1   8   8  16   1 
 Sum 2-char digits   4   1   3   9   1   8   8   7   1 
 ``` 

- 然後將所有數字相加，`4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`。
- 由於 `42` 不是 `10` 的倍數，該數字無效。

在本實驗中，你將使用 Luhn 演算法建置一個信用卡驗證器。

**目標：** 完成以下使用者故事並通過所有測試以完成實驗。

**使用者故事：**

1. 你應該定義一個名為 `verify_card_number` 的函式（程式），該函式接受一個數字字串（表述卡號），並根據 Luhn 演算法驗證其是否有效。

2. 在 `verify_card_number` 函式（程式）內：

   - 你應該控制代碼任何可能出現在傳入卡號中的破折號或空格。
   - 如果卡號有效，傳回 `VALID!`；否則，傳回 `INVALID!`。

當你完成專案時，根據輸入你應該會看到以下訊息：

| Card Number         | Message  |
| ------------------- | -------- |
| 453914889           | VALID!   |
| 4111-1111-1111-1111 | VALID!   |
| 1234 5678 9012 3456 | INVALID! |


# --hints--

你應該有一個名為 `verify_card_number` 的函式（程式）。

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` 應該傳回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` 應該傳回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` 應該傳回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` 應該傳回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

當使用有效的信用卡號碼呼叫 `verify_card_number` 時，應傳回 `VALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

當使用無效的信用卡號碼呼叫 `verify_card_number` 時，應傳回 `INVALID!`。

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
