---
id: 69008059d43d6751385883c1
title: Tafuta makosa ya opereta ya kuongeza na kupunguza katika programu yenye hitilafu
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Rafiki yako amekuomba umsaidie kutafuta makosa katika msimbo wake. Anajifunza kutumia opereta ya kuongeza na kupunguza lakini anapata matokeo yasiyotarajiwa kwenye konsoli.

Zoezi lako ni kumsaidia kubaini makosa na kusahihisha msimbo ili matokeo yanayotokea kwenye konsoli yaendane na yanayotarajiwa.

Unaweza kurejea funzo zilizopita kuhusu opereta ya kuongeza na kupunguza pamoja na noti za kiambishi awali na kiambishi cha mwisho.

**Hadithi za mtumizi:**

1. Unapaswa kuwa na kigezo kinachoitwa `updatedCounter` kinachotoa thamani ya `11` kwa usahihi kwa kutumia opereta ya kuongeza. Sasisha usemi wa kihesabu kutumia noti sahihi ya kiambishi awali au kiambishi cha mwisho badala ya kuweka thamani moja kwa moja.
2. Unapaswa kuwa na kigezo kinachoitwa `finalScore` kinachotoa thamani ya `8` kwa usahihi kwa kutumia opereta ya kuongeza. Sasisha usemi wa kihesabu kutumia noti sahihi ya kiambishi awali au kiambishi cha mwisho badala ya kuweka thamani moja kwa moja.
3. Unapaswa kuwa na kigezo kinachoitwa `updatedCoins` kinachotoa thamani ya `2` kwa usahihi kwa kutumia opereta ya kupunguza. Sasisha usemi wa kihesabu kutumia noti sahihi ya kiambishi awali au kiambishi cha mwisho badala ya kuweka thamani moja kwa moja.
4. Unapaswa kuwa na kigezo kinachoitwa `newHealth` kinachotoa thamani ya `7` kwa usahihi kwa kutumia opereta ya kupunguza. Sasisha usemi wa kihesabu kutumia noti sahihi ya kiambishi awali au kiambishi cha mwisho badala ya kuweka thamani moja kwa moja.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `updatedCounter` kinachotumia noti ya kiambishi awali ya kuongeza (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` inapaswa kuwa sawa na `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Unapaswa kuwa na kigezo kinachoitwa `finalScore` kinachotumia noti ya kiambishi cha mwisho ya kuongeza (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` inapaswa kuwa sawa na `8`.

```js
assert.strictEqual(finalScore, 8);
```

Unapaswa kuwa na kigezo kinachoitwa `updatedCoins` kinachotumia noti ya kiambishi awali ya kupunguza (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` inapaswa kuwa sawa na `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Unapaswa kuwa na kigezo kinachoitwa `newHealth` kinachotumia noti ya kiambishi cha mwisho ya kupunguza (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` inapaswa kuwa sawa na `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
