---
id: 68dbc6b7e8f794a2ae69bc74
title: تنفيذ خوارزمية N-Queens
challengeType: 27
dashedName: implement-the-n-queens-algorithm
---

# --description--

تطلب مشكلة N-Queens وضع N ملكة على رقعة شطرنج بحجم N×N بحيث لا تهاجم أي ملكتين بعضهما البعض (لا يتشاركان في صف أو عمود أو قطر).

على سبيل المثال، إذا كانت الرقعة 4x4، فإن أحد الترتيبات الصحيحة هو:

```md
[1, 3, 0, 2]
```

هذا يعني أنه في الصف 0، توضع الملكة في العمود 1؛ في الصف 1، توضع الملكة في العمود 3؛ في الصف 2، توضع الملكة في العمود 0؛ وفي الصف 3، توضع الملكة في العمود 2.

بصريًا، يبدو هذا الترتيب كما يلي:

```md
. Q . .
. . . Q
Q . . .
. . Q .
```

حيث تمثل `Q` ملكة وتمثل `.` مربعًا فارغًا.

في هذا المختبر، ستنفذ حل مشكلة N-Queens باستخدام طريقة البحث بالعمق أولاً.


**الهدف**: أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تحتوي على دالة باسم `dfs_n_queens`.
2. يجب أن تقبل الدالة معلمة واحدة فقط: عدد صحيح `n`.
3. إذا كان `n` أقل من `1`، يجب أن تُرجع الدالة قائمة فارغة (`[]`).
4. يجب أن تُرجع الدالة قائمة من الحلول؛ كل حل هو قائمة بطول `n`، حيث العنصر في الفهرس `i` هو فهرس العمود (يبدأ من الصفر) للملكة في الصف `i`.

# --hints--

يجب أن يكون لديك دالة باسم `dfs_n_queens` تأخذ معلمة واحدة.

```js
({ test: () => runPython(`
  import inspect
  assert inspect.isfunction(dfs_n_queens)
  sig = inspect.signature(dfs_n_queens)
  assert len(sig.parameters) == 1
`) })
```

إذا كان `n` أقل من `1`، يجب أن تُرجع الدالة قائمة فارغة.

```js
({ test: () => runPython(`
  assert dfs_n_queens(0) == []
  assert dfs_n_queens(-1) == []
  assert dfs_n_queens(-5) == []
`) })
```

يجب أن تُرجع الدالة قائمة من الحلول، حيث كل حل هو قائمة بطول `n`.

```js
({ test: () => runPython(`
  result = dfs_n_queens(4)
  assert isinstance(result, list)
  for solution in result:
    assert isinstance(solution, list)
    assert len(solution) == 4
`) })
```

`dfs_n_queens(1)` يجب أن تُرجع `[[0]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(1) == [[0]]
`) })
```

`dfs_n_queens(2)` يجب أن تُرجع `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(2) == []
`) })
```

`dfs_n_queens(3)` يجب أن تُرجع `[]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(3) == []
`) })
```

`dfs_n_queens(4)` يجب أن تُرجع `[[1, 3, 0, 2], [2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(4) == [[1, 3, 0, 2], [2, 0, 3, 1]]
`) })
```

`dfs_n_queens(5)` يجب أن تُرجع `[[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]`.

```js
({ test: () => runPython(`
  assert dfs_n_queens(5) == [[0, 2, 4, 1, 3], [0, 3, 1, 4, 2], [1, 3, 0, 2, 4], [1, 4, 2, 0, 3], [2, 0, 3, 1, 4], [2, 4, 1, 3, 0], [3, 0, 2, 4, 1], [3, 1, 4, 2, 0], [4, 1, 3, 0, 2], [4, 2, 0, 3, 1]]
`) })
```

`len(dfs_n_queens(5))` يجب أن تكون `10`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(5)) == 10
`) })
```

`len(dfs_n_queens(8))` يجب أن تكون `92`.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(8)) == 92
`) })
```

`dfs_n_queens` يجب أن تُرجع النتيجة الصحيحة.

```js
({ test: () => runPython(`
  assert len(dfs_n_queens(7)) == 40
  assert dfs_n_queens(6) == [[1, 3, 5, 0, 2, 4], [2, 5, 1, 4, 0, 3], [3, 0, 4, 1, 5, 2], [4, 2, 0, 5, 3, 1]]
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def dfs_n_queens(n):
    if n < 1:
        return []

    results = []

    def is_safe(queens, row, col):
        for r, c in enumerate(queens):
            if c == col or abs(row - r) == abs(col - c):
                return False
        return True

    def dfs(queens):
        row = len(queens)
        if row == n:
            results.append(queens[:])
            return
        for col in range(n):
            if is_safe(queens, row, col):
                queens.append(col)
                dfs(queens)
                queens.pop()

    dfs([])
    return results
```
