---
id: 6823e637c1c0ed56f781b4fc
title: Paso 11
challengeType: 0
dashedName: step-11
---

# --description--

Finalmente, JavaScript hace que tu página web sea interactiva; te permite indicar a la página qué hacer cuando alguien hace clic en un botón, envía un formulario, o muchas otras acciones.

Para el último paso del taller, agrega otro elemento `h3` y `p` a la página describiendo JavaScript.

Primero, agrega un elemento `h3` con el texto:

```md
Introduction to JavaScript
```

Luego, debajo de ese elemento `h3`, agrega un elemento `p` con el siguiente texto:

```md
JavaScript adds interactivity to a webpage
```

# --hints--

Deberías agregar un tercer elemento `h3` a la página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 3);
```

El texto del elemento `h3` debería ser `Introduction to JavaScript`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*JavaScript\s*\<\/h3\s*\>/i);
```

Deberías agregar un cuarto elemento `p` a la página.

```js
assert.lengthOf(document.querySelectorAll("p"), 4);
```

El elemento `p` debería tener el texto `JavaScript adds interactivity to a webpage`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*JavaScript\s*adds\s*interactivity\s*to\s*a\s*web\s*page\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

<h3>Introduction to CSS</h3>
<p>CSS is used to style a webpage</p>

<h3>Introduction to JavaScript</h3>
<p>JavaScript adds interactivity to a webpage</p>
```
