---
id: 6823d9ac8bdc3853df65a1ff
title: Paso 7
challengeType: 0
dashedName: step-7
---

# --description--

Hay seis elementos de encabezado en HTML: `h1` a `h6`. Se utilizan para mostrar la importancia de las secciones en la página web, con `h1` siendo el más importante y `h6` el menos importante.

Debajo de su elemento `p`, agregue un encabezado `h3` con el texto:

```md
Introduction to HTML
```

# --hints--

Tu elemento `h3` debe tener una etiqueta de apertura `<h3>`.

```js
assert.exists(document.querySelector("h3"));
```

Tu elemento `h3` debe tener una etiqueta de cierre `</h3>`.

```js
assert.match(code, /<\/h3\s*\>/);
```

Su elemento `h3` debe tener el texto `Introduction to HTML`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*HTML\s*\<\/h3\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full-Stack Curriculum</h2>
<p>Learn the skills to become a full-stack developer</p>

--fcc-editable-region--

--fcc-editable-region--
```
