---
id: 682ba2318000b62f179bdf04
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

HTML está compuesto por elementos. El primero que utilizarás es el elemento `h1`:

```html
<h1>Welcome to freeCodeCamp</h1>
```

Comienza con una etiqueta de apertura (`<h1>`), termina con una etiqueta de cierre (`</h1>`) y tiene el texto que mostrará entre las etiquetas.

Convierte tu texto `Welcome to freeCodeCamp` en un elemento `h1` agregando una etiqueta de apertura antes de él y una etiqueta de cierre después.

# --hints--

Tu elemento `h1` debe tener una etiqueta de apertura `<h1>`.

```js
assert.exists(document.querySelector("h1"));
```

Tu elemento `h1` debe tener una etiqueta de cierre `</h1>`.

```js
assert.match(code, /<\/h1\s*\>/);
```

Tu elemento `h1` debería verse así: `<h1>Welcome to freeCodeCamp</h1>`.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h1\s*\>\s*Welcome\s*to\s*freeCodeCamp\s*\<\/h1\s*\>/i);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
Welcome to freeCodeCamp
--fcc-editable-region--
```
