---
id: 682cd206883fc7b25eb539c5
title: Paso 3
challengeType: 0
dashedName: step-3
---

# --description--

Observe que el HTML que escribe en el editor aparece en la vista previa. En este taller, escribirá el HTML para una página web parcial del currículo.

Debajo de su `h1` elemento, escriba lo siguiente en la línea vacía:

```md
Full-Stack Curriculum
```

# --hints--

Debes tener el texto `Full-Stack Curriculum` en tu editor. Verifica la ortografía.

```js
// purposefully providing forgiveness in spacing and casing to remove some friction and
// keep camper retention this early on.
assert.match(code, /full-stack\s*curriculum/i)
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
--fcc-editable-region--

--fcc-editable-region--
```
