---
id: 683921c4769fd23dbadec2fe
title: Paso 10
challengeType: 0
dashedName: step-10
---

# --description--

Mientras que HTML define la estructura y el contenido de una página web, CSS se utiliza para agregar estilo: cosas como colores, fuentes, espaciado y diseño.

Debajo del `h3` que acaba de agregar, agregue otro elemento de párrafo con el texto:

```md
CSS is used to style a webpage
```

# --hints--

Deberías añadir un segundo elemento `h3` a la página.

```js
assert.lengthOf(document.querySelectorAll("h3"), 2);
```

El texto de tu elemento `h3` debería ser `Introduction to CSS`. 

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<h3\s*\>\s*Introduction\s*to\s*CSS\s*\<\/h3\s*\>/i);
```

Debes agregar un tercer elemento `p` a la página.

```js
assert.lengthOf(document.querySelectorAll("p"), 3);
```

Tu elemento `p` debe tener el texto `CSS is used to style a webpage`. Puede copiar el texto de las instrucciones.

```js
// purposefully removing friction for early users to help improve retention in early lessons
// this if very forgiving of spaces and casing
assert.match(code, /\<p\s*\>\s*CSS\s*is\s*used\s*to\s*style\s*a\s*webpage\s*\<\/p\s*\>/i);
```

# --seed--

## --seed-contents--

```html
<h1>Welcome to freeCodeCamp</h1>
<h2>Full Stack Curriculum</h2>
<p>Learn the skills to become a full stack developer</p>

<h3>Introduction to HTML</h3>
<p>HTML represents the content and structure of a webpage</p>

--fcc-editable-region--
<h3>Introduction to CSS</h3>

--fcc-editable-region--
```
