---
id: 67d83df6f82eda3868dd2a84
title: 创建一个 RPG 字符
challengeType: 27
dashedName: build-an-rpg-character
---

# --description--

在本实验中，你将通过创建一个用于 RPG 冒险的字符的小型应用来练习 Python 的基础知识。

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户故事：**

1. 你应该有一个名为 `create_character` 的函数。
1. 该函数应按顺序接受一个字符名称，后跟三个属性：力量、智力和魅力。
1. 字符名称应被验证：
   - 如果字符名称不是字串，函数应返回 `The character name should be a string`。
   - 如果字符名称是空字串，函数应返回 `The character should have a name`。
   - 如果字符名称超过 10 个字符，函数应返回 `The character name is too long`。
   - 如果字符名称包含空格，函数应返回 `The character name should not contain spaces`。
1. 这些统计数据也应该被验证：
   - 如果一个或多个属性不是整数，函数应返回 `All stats should be integers`。
   - 如果一个或多个属性小于 1，函数应返回 `All stats should be no less than 1`。
   - 如果一个或多个属性超过 4，函数应返回 `All stats should be no more than 4`。
   - 如果所有属性的总和不等于 7，函数应返回 `The character should start with 7 points`。
1. 如果所有值都通过验证，函数应返回一个包含四行的字串：
   - 第一行应包含字符名称
   - 第 2-4 行应以状态缩略语 `STR`、`INT` 或 `CHA`（按此顺序）开头，后跟一个空格，然后是与状态值相等数量的实心点（`●`），以及补足到 10 个的空心点（`○`）。示例：如果力量值为 3，则应有 3 个实心点，后跟 7 个空心点。点符号由编辑器提供。

这是 `create_character('ren', 4, 2, 1)` 应该返回的字串：

   ```md
   ren
   STR ●●●●○○○○○○
   INT ●●○○○○○○○○
   CHA ●○○○○○○○○○
   ```

注意：虽然 `str` 和 `int` 是统计数据的常见缩略语，但请记住它们是 Python 中的保留关键字，不应作为变量名使用。

# --hints--

你应该有一个名为 `create_character` 的函数。

```js
({
    test: () => runPython(
`
assert _Node(_code).has_function('create_character')
`
    )
})
```

当调用 `create_character` 且第一个参数不是字串时，应返回 `The character name should be a string`。

```js
({
    test: () => runPython(
        `
wrong_types = [3, 5, {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}]
for wrong_type in wrong_types:
    assert create_character(wrong_type, 1, 2, 4) == 'The character name should be a string'
`
    )
})
```

当调用 `create_character` 且第一个参数是空字串时，它应返回 `The character should have a name`。

```js
({
    test: () => runPython(
        `
assert create_character('', 4, 2, 1) == 'The character should have a name'
`
    )
})
```

当调用 `create_character` 且第一个参数长度超过 10 个字符时，应返回 `The character name is too long`。

```js
({
    test: () => runPython(
        `
for length in range(11, 27):
    assert create_character('abcdefghijklmnopqrstuvwxyz'[:length], 4, 2, 1) == 'The character name is too long'
`
    )
})
```

当 `character` 不超过 10 个字符时，`create_character` 函数不应提示字符过长。

```js
({
    test: () => runPython(
        `
for length in range(0, 11):
  assert create_character('aaabbcccdd'[:length], 4, 2, 1) != 'The character name is too long'
`
    )
})
```

当 `create_character` 被调用且第一个参数包含空格时，应返回 `The character name should not contain spaces`。

```js
({
    test: () => runPython(
        `
assert create_character('cha cha', 4, 1, 2) == 'The character name should not contain spaces'
`
    )
})
```

当 `create_character` 被调用且第二、第三或第四个参数不是整数时，应返回 `All stats should be integers`。

```js
({
    test: () => runPython(
        `
wrong_types = ['3', '5', {'a': 3}, 3.1, [1, 2], (1, 2), {1, 2}, 'str', 'friend']
for wrong_type in wrong_types:
    assert create_character('friend', wrong_type, 2, 1) == 'All stats should be integers'
    assert create_character('friend', 2, wrong_type, 1) == 'All stats should be integers'
    assert create_character('friend', 2, 1, wrong_type) == 'All stats should be integers'
`
    )
})
```

当 `create_character` 被调用且第二、第三或第四个参数小于 `1` 时，应返回 `All stats should be no less than 1`。

```js
({
    test: () => runPython(
        `
expected = 'All stats should be no less than 1'

assert create_character('ren', 0, 4, 3) == expected
assert create_character('ren', 4, 0, 3) == expected
assert create_character('ren', 4, 3, 0) == expected
assert create_character('ren', -1, 4, 4) == expected
assert create_character('ren', 4, -1, 4) == expected
assert create_character('ren', 4, 4, -1) == expected
`
    )
})
```

当 `create_character` 被调用且第二、第三或第四个参数大于 `4` 时，它应返回 `All stats should be no more than 4`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 1, 5, 1) == 'All stats should be no more than 4'
assert create_character('ren', 5, 1, 1) == 'All stats should be no more than 4'
assert create_character('ren', 1, 1, 5) == 'All stats should be no more than 4'
`
    )
})
```

当调用 `create_character` 并且第二、第三或第四个参数之和不等于 `7` 时，应返回 `The character should start with 7 points`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 4, 4) == 'The character should start with 7 points'
assert create_character('ren', 1, 1, 1) == 'The character should start with 7 points'
assert create_character('ren', 1, 2, 3) == 'The character should start with 7 points'
`
    )
})
```

`create_character('ren', 4, 2, 1)` 应该返回 `ren\nSTR ●●●●○○○○○○\nINT ●●○○○○○○○○\nCHA ●○○○○○○○○○`。

```js
({
    test: () => runPython(
        `
assert create_character('ren', 4, 2, 1) == f'ren\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})
```

当使用有效值调用 `create_character` 时，应按要求输出角色属性。

```js
({
    test: () => runPython(
        `
assert create_character('Bob', 1, 2, 4) == f'Bob\\nSTR ●○○○○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●●●●○○○○○○'
assert create_character('aaabbbcccd', 4, 2, 1) == f'aaabbbcccd\\nSTR ●●●●○○○○○○\\nINT ●●○○○○○○○○\\nCHA ●○○○○○○○○○'
`
    )
})

```

# --seed--

## --seed-contents--

```py
full_dot = '●'
empty_dot = '○'


```

# --solutions--

```py
full_dot = '●'
empty_dot = '○'

def create_character(character_name, strength, intelligence, charisma):
    if type(character_name) != str:
        return 'The character name should be a string'
    if len(character_name) == 0:
        return 'The character should have a name'
    if len(character_name) > 10:
        return 'The character name is too long'
    if ' ' in character_name:
        return 'The character name should not contain spaces'

    if type(strength) != int or type(intelligence) != int or type(charisma) != int:
       return 'All stats should be integers'
    if strength < 1 or intelligence < 1 or charisma < 1:
        return 'All stats should be no less than 1'
    if strength > 4 or intelligence > 4 or charisma > 4:
        return 'All stats should be no more than 4'
    if strength + intelligence + charisma != 7:
        return 'The character should start with 7 points'

    character_sheet = f'{character_name}\n'
    character_sheet += f'STR {full_dot * strength}{empty_dot * (10 - strength)}\n'
    character_sheet += f'INT {full_dot * intelligence}{empty_dot * (10 - intelligence)}\n'
    character_sheet += f'CHA {full_dot * charisma}{empty_dot * (10 - charisma)}'
    return character_sheet
```
