---
id: 6814d8e1516e86b171929de4
title: "Завдання 1: баланс голосних"
challengeType: 28
dashedName: challenge-1
---

# --description--

Дано рядок. Визначте, чи дорівнює кількість голосних у першій половині рядка кількості голосних у другій половині.

- Рядок може містити будь-які символи.
- Літери `a`, `e`, `i`, `o` і `u`, у верхньому або нижньому регістрі, вважаються голосними.
- Якщо кількість символів у рядку непарна, ігноруйте центральний символ.

# --hints--

`isBalanced("racecar")` має повернути `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` має повернути `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` має повернути `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` має повернути `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` має повернути `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` має повернути `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` має повернути `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

`isBalanced("a")` має повернути `true`.

```js
assert.isTrue(isBalanced("a"));
```

`isBalanced("E")` має повернути `true`.

```js
assert.isTrue(isBalanced("E"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const chars = [...s];
  const half = Math.floor(chars.length / 2);

  const firstHalf = chars.slice(0, half);
  const secondHalf =
    chars.length % 2 === 0 ? chars.slice(half) : chars.slice(half + 1);

  const countVowels = arr =>
    arr.filter(c => vowels.includes(c.toLowerCase())).length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
