---
id: 6814d8e1516e86b171929de4
title: "Завдання 1: баланс голосних"
challengeType: 28
dashedName: challenge-1
---

# --description--

Дано рядок, визначте, чи дорівнює кількість голосних у першій половині рядка кількості голосних у другій половині.

- Рядок може містити будь-які символи.
- Літери `a`, `e`, `i`, `o` та `u`, у верхньому або нижньому регістрі, вважаються голосними.
- Якщо кількість символів у рядку непарна, ігноруйте центральний символ.

# --hints--

`isBalanced("racecar")` має повернути `true`.

```js
assert.isTrue(isBalanced("racecar"));
```

`isBalanced("Lorem Ipsum")` має повернути `true`.

```js
assert.isTrue(isBalanced("Lorem Ipsum"));
```

`isBalanced("Kitty Ipsum")` має повернути `false`.

```js
assert.isFalse(isBalanced("Kitty Ipsum"));
```

`isBalanced("string")` має повернути `false`.

```js
assert.isFalse(isBalanced("string"));
```

`isBalanced(" ")` має повернути `true`.

```js
assert.isTrue(isBalanced(" "));
```

`isBalanced("abcdefghijklmnopqrstuvwxyz")` має повернути `false`.

```js
assert.isFalse(isBalanced("abcdefghijklmnopqrstuvwxyz"));
```

`isBalanced("123A#b!E&*456-o.U")` має повернути `true`.

```js
assert.isTrue(isBalanced("123A#b!E&*456-o.U"));
```

# --seed--

## --seed-contents--

```js
function isBalanced(s) {

  return s;
}
```

# --solutions--

```js
function isBalanced(s) {
  const vowels = 'aeiou';
  const half = Math.floor(s.length / 2);

  let firstHalf = s.slice(0, half);
  let secondHalf = s.length % 2 === 0 ? s.slice(half) : s.slice(half + 1);

  const countVowels = str =>
    str
      .toLowerCase()
      .split('')
      .filter(c => vowels.includes(c))
      .length;

  return countVowels(firstHalf) === countVowels(secondHalf);
}
```
