---
id: 681cb1afdab50c87ddb2e517
title: "Завдання 6: перевірка анаграми"
challengeType: 28
dashedName: challenge-6
---

# --description--

Дано два рядки, визначте, чи є вони анаграмами один одного (містять ті самі символи в будь-якому порядку).

- Ігноруйте регістр і порожній простір.

# --hints--

`areAnagrams("listen", "silent")` має повернути `true`.

```js
assert.isTrue(areAnagrams("listen", "silent"));
```

`areAnagrams("School master", "The classroom")` має повернути `true`.

```js
assert.isTrue(areAnagrams("School master", "The classroom"));
```

`areAnagrams("A gentleman", "Elegant man")` має повернути `true`.

```js
assert.isTrue(areAnagrams("A gentleman", "Elegant man"));
```

`areAnagrams("Hello", "World")` має повернути `false`.

```js
assert.isFalse(areAnagrams("Hello", "World"));
```

`areAnagrams("apple", "banana")` має повернути `false`.

```js
assert.isFalse(areAnagrams("apple", "banana"));
```

`areAnagrams("cat", "dog")` має повернути `false`.

```js
assert.isFalse(areAnagrams("cat", "dog"));
```

# --seed--

## --seed-contents--

```js
function areAnagrams(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function areAnagrams(str1, str2) {
  const clean = (str) =>
    str.replace(/\s+/g, '').toLowerCase().split('').sort().join('');
    
  return clean(str1) === clean(str2);
}
```
