---
id: 681cb1afdab50c87ddb2e516
title: "Завдання 5: jbelmud текст"
challengeType: 28
dashedName: challenge-5
---

# --description--

Дано рядок, поверніть перемішану версію цього рядка, де кожне слово трансформується за такими правилами:

- Перша і остання букви слова залишаються на своїх місцях
- Всі букви між першою та останньою сортуються в алфавітному порядку.
- Вхідні рядки не міститимуть розділових знаків і будуть повністю в нижньому регістрі.

# --hints--

`jbelmu("hello world")` має повертати `"hello wlord"`.

```js
assert.equal(jbelmu("hello world"), "hello wlord");
```

`jbelmu("i love jumbled text")` має повертати `"i love jbelmud text"`.

```js
assert.equal(jbelmu("i love jumbled text"), "i love jbelmud text");
```

`jbelmu("freecodecamp is my favorite place to learn to code")` має повертати `"faccdeeemorp is my faiortve pacle to laern to cdoe"`.

```js
assert.equal(jbelmu("freecodecamp is my favorite place to learn to code"), "faccdeeemorp is my faiortve pacle to laern to cdoe");
```

`jbelmu("the quick brown fox jumps over the lazy dog")` має повертати `"the qciuk borwn fox jmpus oevr the lazy dog"`.

```js
assert.equal(jbelmu("the quick brown fox jumps over the lazy dog"), "the qciuk borwn fox jmpus oevr the lazy dog");
```

# --seed--

## --seed-contents--

```js
function jbelmu(text) {

  return text;
}
```

# --solutions--

```js
function jbelmu(text) {
  return text
    .split(' ')
    .map((word) => {
      if (word.length <= 3) return word;
      const first = word[0];
      const last = word[word.length - 1];
      const middle = word
        .slice(1, -1)
        .split('')
        .sort()
        .join('');
      return first + middle + last;
    })
    .join(' ');
}
```
