---
id: 681cb1afdab50c87ddb2e515
title: "Завдання 4: S  P  A  C  E  J  A  M"
challengeType: 28
dashedName: challenge-4
---

# --description--

Дано рядок. Видаліть усі пробіли з рядка, вставте по два пробіли між кожним символом, перетворіть усі літери алфавіту на верхній регістр і поверніть результат.

- Неалфавітні символи мають залишатися без змін (окрім пробілів).

# --hints--

`spaceJam("freeCodeCamp")` має повернути `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("freeCodeCamp"), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("   free   Code   Camp   ")` має повернути `"F  R  E  E  C  O  D  E  C  A  M  P"`.

```js
assert.equal(spaceJam("   free   Code   Camp   "), "F  R  E  E  C  O  D  E  C  A  M  P");
```

`spaceJam("Hello World?!")` має повернути `"H  E  L  L  O  W  O  R  L  D  ?  !"`.

```js
assert.equal(spaceJam("Hello World?!"), "H  E  L  L  O  W  O  R  L  D  ?  !");
```

`spaceJam("C@t$ & D0g$")` має повернути `"C  @  T  $  &  D  0  G  $"`.

```js
assert.equal(spaceJam("C@t$ & D0g$"), "C  @  T  $  &  D  0  G  $");
```

`spaceJam("allyourbase")` має повернути `"A  L  L  Y  O  U  R  B  A  S  E"`.

```js
assert.equal(spaceJam("all your base"), "A  L  L  Y  O  U  R  B  A  S  E");
```

# --seed--

## --seed-contents--

```js
function spaceJam(s) {

  return s;
}
```

# --solutions--

```js
function spaceJam(s) {
  return s.toUpperCase().replace(/\s+/g, '').split('').join('  ');
}
```
