---
id: 681cb1b0dab50c87ddb2e51b
title: "Завдання 10: 3 страйки"
challengeType: 28
dashedName: challenge-10
---

# --description--

Дано ціле число від 1 до 10 000. Поверніть кількість чисел від 1 до цього числа, квадрат яких містить принаймні одну цифру 3.

# --hints--

`squaresWithThree(1)` має повертати `0`.

```js
assert.equal(squaresWithThree(1), 0);
```

`squaresWithThree(10)` має повертати `1`.

```js
assert.equal(squaresWithThree(10), 1);
```

`squaresWithThree(100)` має повертати `19`.

```js
assert.equal(squaresWithThree(100), 19);
```

`squaresWithThree(1000)` має повертати `326`.

```js
assert.equal(squaresWithThree(1000), 326);
```

`squaresWithThree(10000)` має повертати `4531`.

```js
assert.equal(squaresWithThree(10000), 4531);
```

# --seed--

## --seed-contents--

```js
function squaresWithThree(n) {

  return n;
}
```

# --solutions--

```js
function squaresWithThree(n) {
  let count = 0;

  for (let i = 1; i <= n; i++) {
    const square = i * i;
    if (square.toString().includes('3')) {
      count++;
    }
  }

  return count;
}
```
