---
id: 681cb1b0dab50c87ddb2e518
title: "Завдання 7: цільова сума"
challengeType: 28
dashedName: challenge-7
---

# --description--

Дано масив чисел і ціле число ціль. Знайдіть два унікальні числа в масиві, які у сумі дорівнюють цільовому значенню. Поверніть масив з індексами цих двох чисел або `"Target not found"`, якщо жодні два числа не дають у сумі ціль.

- Повернутий масив повинен містити індекси у зростаючому порядку.

# --hints--

`findTarget([2, 7, 11, 15], 9)` має повертати `[0, 1]`.

```js
assert.deepEqual(findTarget([2, 7, 11, 15], 9), [0, 1]);
```

`findTarget([3, 2, 4, 5], 6)` має повертати `[1, 2]`.

```js
assert.deepEqual(findTarget([3, 2, 4, 5], 6), [1, 2]);
```

`findTarget([1, 3, 5, 6, 7, 8], 15)` має повертати `[4, 5]`.

```js
assert.deepEqual(findTarget([1, 3, 5, 6, 7, 8], 15), [4, 5]);
```

`findTarget([1, 3, 5, 7], 14)` має повертати `"Target not found"`.

```js
assert.equal(findTarget([1, 3, 5, 7], 14), "Target not found");
```

# --seed--

## --seed-contents--

```js
function findTarget(arr, target) {

  return arr;
}
```

# --solutions--

```js
function findTarget(arr, target) {
  for (let i = 0; i < arr.length; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] + arr[j] === target) {
        return [i, j];
      }
    }
  }
  return 'Target not found'
}
```
