---
id: 681cb1b0dab50c87ddb2e51a
title: "Завдання 9: сума квадратів"
challengeType: 28
dashedName: challenge-9
---

# --description--

Дано додатне ціле число до 1 000. Поверніть суму квадратів усіх цілих чисел від 1 до цього числа.

# --hints--

`sumOfSquares(5)` має повернути `55`.

```js
assert.equal(sumOfSquares(5), 55);
```

`sumOfSquares(10)` має повернути `385`.

```js
assert.equal(sumOfSquares(10), 385);
```

`sumOfSquares(25)` має повернути `5525`.

```js
assert.equal(sumOfSquares(25), 5525);
```

`sumOfSquares(500)` має повернути `41791750`.

```js
assert.equal(sumOfSquares(500), 41791750);
```

`sumOfSquares(1000)` має повернути `333833500`.

```js
assert.equal(sumOfSquares(1000), 333833500);
```

# --seed--

## --seed-contents--

```js
function sumOfSquares(n) {

  return n;
}
```

# --solutions--

```js
function sumOfSquares(n) {
  let total = 1;
  for(let i = 2; i <= n; i++) {
      total += i * i;
  }
  return total;
}
```
