---
id: 681cb05adab50c87ddb2e513
title: "Завдання 2: перевірка основи"
challengeType: 28
dashedName: challenge-2
---

# --description--

Дано рядок, який представляє число, та ціле число основи від 2 до 36. Визначте, чи є це число дійсним у цій основі.

- Рядок може містити цілі числа, а також символи верхнього або нижнього регістру.
- Перевірка має бути нечутливою до регістру.
- Основа може бути будь-яким числом від 2 до 36.
- Число є дійсним, якщо кожен символ є дійсною цифрою у заданій основі.
- Приклади дійсних цифр для основ:
  - Основа 2: 0-1
  - Основа 8: 0-7
  - Основа 10: 0-9
  - Основа 16: 0-9 та A-F
  - Основа 36: 0-9 та A-Z

# --hints--

`isValidNumber("10101", 2)` має повернути `true`.

```js
assert.isTrue(isValidNumber("10101", 2))
```

`isValidNumber("10201", 2)` має повернути `false`.

```js
assert.isFalse(isValidNumber("10201", 2))
```

`isValidNumber("76543210", 8)` має повернути `true`.

```js
assert.isTrue(isValidNumber("76543210", 8))
```

`isValidNumber("9876543210", 8)` має повернути `false`.

```js
assert.isFalse(isValidNumber("9876543210", 8))
```

`isValidNumber("9876543210", 10)` має повернути `true`.

```js
assert.isTrue(isValidNumber("9876543210", 10))
```

`isValidNumber("ABC", 10)` має повернути `false`.

```js
assert.isFalse(isValidNumber("ABC", 10))
```

`isValidNumber("ABC", 16)` має повернути `true`.

```js
assert.isTrue(isValidNumber("ABC", 16))
```

`isValidNumber("Z", 36)` має повернути `true`.

```js
assert.isTrue(isValidNumber("Z", 36))
```

`isValidNumber("ABC", 20)` має повернути `true`.

```js
assert.isTrue(isValidNumber("ABC", 20))
```

`isValidNumber("4B4BA9", 16)` має повернути `true`.

```js
assert.isTrue(isValidNumber("4B4BA9", 16))
```

`isValidNumber("5G3F8F", 16)` має повернути `false`.

```js
assert.isFalse(isValidNumber("5G3F8F", 16))
```

`isValidNumber("5G3F8F", 17)` має повернути `true`.

```js
assert.isTrue(isValidNumber("5G3F8F", 17))
```

`isValidNumber("abc", 10)` має повернути `false`.

```js
assert.isFalse(isValidNumber("abc", 10))
```

`isValidNumber("abc", 16)` має повернути `true`.

```js
assert.isTrue(isValidNumber("abc", 16))
```

`isValidNumber("AbC", 16)` має повернути `true`.

```js
assert.isTrue(isValidNumber("AbC", 16))
```

`isValidNumber("z", 36)` має повернути `true`.

```js
assert.isTrue(isValidNumber("z", 36))
```

# --seed--

## --seed-contents--

```js
function isValidNumber(n, base) {

  return n;
}
```

# --solutions--

```js
function isValidNumber(n, base) {
  return new RegExp(`^[${'0123456789abcdefghijklmnopqrstuvwxyz'.substring(0, base)}]+\$`, "i").test(n);
}
```
