---
id: 681cb1b0dab50c87ddb2e519
title: "Завдання 8: факторіал"
challengeType: 28
dashedName: challenge-8
---

# --description--

Дано ціле число від нуля до 20. Потрібно повернути факторіал цього числа. Факторіал числа — це добуток усіх чисел від 1 до заданого числа.

- Факторіал нуля дорівнює 1.

# --hints--

`factorial(0)` має повертати `1`.

```js
assert.equal(factorial(0), 1);
```

`factorial(5)` має повертати `120`.

```js
assert.equal(factorial(5), 120);
```

`factorial(20)` має повертати `2432902008176640000`.

```js
assert.equal(factorial(20), 2432902008176640000);
```

# --seed--

## --seed-contents--

```js
function factorial(n) {

  return n;
}
```

# --solutions--

```js
function factorial(n) {
  return n == 0 ? 1 : n * factorial(n - 1);
}
```
