---
id: 681cb1a2dab50c87ddb2e514
title: "Завдання 3: послідовність Фібоначчі"
challengeType: 28
dashedName: challenge-3
---

# --description--

Послідовність Фібоначчі — це ряд чисел, де кожне число є сумою двох попередніх. Починаючи з `0` і `1`, перші 10 чисел послідовності — це `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

Дано масив, що містить перші два числа послідовності Фібоначчі, та ціле число, яке представляє довжину послідовності. Поверніть масив, що містить послідовність заданої довжини.

- Ваша функція має працювати з послідовностями будь-якої довжини, більшої або рівної нулю.
- Якщо довжина дорівнює нулю, поверніть порожній масив.
- Зверніть увагу, що початкові числа є частиною послідовності.

# --hints--

`fibonacciSequence([0, 1], 20)` має повернути `[0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 20), [0, 1, 1, 2, 3, 5, 8, 13, 21, 34, 55, 89, 144, 233, 377, 610, 987, 1597, 2584, 4181]);
```

`fibonacciSequence([21, 32], 1)` має повернути `[21]`.

```js
assert.deepEqual(fibonacciSequence([21, 32], 1), [21]);
```

`fibonacciSequence([0, 1], 0)` має повернути `[]`.

```js
assert.deepEqual(fibonacciSequence([0, 1], 0), []);
```

`fibonacciSequence([10, 20], 2)` має повернути `[10, 20]`.

```js
assert.deepEqual(fibonacciSequence([10, 20], 2), [10, 20]);
```

`fibonacciSequence([123456789, 987654321], 5)` має повернути `[123456789, 987654321, 1111111110, 2098765431, 3209876541]`.

```js
assert.deepEqual(fibonacciSequence([123456789, 987654321], 5), [123456789, 987654321, 1111111110, 2098765431, 3209876541]);
```

# --seed--

## --seed-contents--

```js
function fibonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function fibonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2];
    sequence.push(nextValue);
  }
  return sequence;
}
```
