---
id: 6821ebce237de8297eaee790
title: "Завдання 14: битва символів"
challengeType: 28
dashedName: challenge-14
---

# --description--

Дано два рядки, які представляють вашу армію та армію супротивника, кожен символ вашої армії бореться з символом на тій самій позиції армії супротивника за такими правилами:

- Символи `a-z` мають відповідно силу `1-26`.
- Символи `A-Z` мають відповідно силу `27-52`.
- Цифри `0-9` мають силу, що дорівнює їхньому значенню.
- Всі інші символи мають значення нуль.
- Кожен символ може брати участь лише в одній битві.

Для кожної битви перемагає сильніший символ. Армія з більшою кількістю перемог виграє війну. Поверніть такі значення:

- `"Opponent retreated"`, якщо у вашій армії більше символів, ніж у армії супротивника.
- `"We retreated"`, якщо у армії супротивника більше символів, ніж у вашій.
- `"We won"`, якщо ваша армія виграла більше битв.
- `"We lost"`, якщо армія супротивника виграла більше битв.
- `"It was a tie"`, якщо обидві армії виграли однакову кількість битв.

# --hints--

`battle("Hello", "World")` має повернути `"We lost"`.

```js
assert.equal(battle("Hello", "World"), "We lost");
```

`battle("pizza", "salad")` має повернути `"We won"`.

```js
assert.equal(battle("pizza", "salad"), "We won");
```

`battle("C@T5", "D0G$")` має повернути `"We won"`.

```js
assert.equal(battle("C@T5", "D0G$"), "We won");
```

`battle("kn!ght", "orc")` має повернути `"Opponent retreated"`.

```js
assert.equal(battle("kn!ght", "orc"), "Opponent retreated");
```

`battle("PC", "Mac")` має повернути `"We retreated"`.

```js
assert.equal(battle("PC", "Mac"), "We retreated");
```

`battle("Wizards", "Dragons")` має повернути `"It was a tie"`.

```js
assert.equal(battle("Wizards", "Dragons"), "It was a tie");
```

`battle("Mr. Smith", "Dr. Jones")` має повернути `"It was a tie"`.

```js
assert.equal(battle("Mr. Smith", "Dr. Jones"), "It was a tie");
```

# --seed--

## --seed-contents--

```js
function battle(myArmy, opposingArmy) {

  return myArmy;
}
```

# --solutions--

```js
function getStrength(soldier) {
  const soldiers = 'abcdefghijklmnopqrstuvwxyzABCDEFGHIJKLMNOPQRSTUVWXYZ';
  let strength = 0;

  if (/\d/.test(soldier)) {
    strength = parseInt(soldier);
  } else if (soldiers.includes(soldier)) {
    strength = soldiers.indexOf(soldier) + 1;
  }

  return strength;
}

function battle(myArmy, opposingArmy) {
  if (myArmy.length > opposingArmy.length) return 'Opponent retreated';
  if (opposingArmy.length > myArmy.length) return 'We retreated';

  let myWins = 0;
  let theirWins = 0;

  for (let i = 0; i < myArmy.length; i++) {
    const mySoldier = myArmy[i];
    const theirSoldier = opposingArmy[i];

    const myStrength = getStrength(mySoldier);
    const theirStrength = getStrength(theirSoldier);

    if (myStrength > theirStrength) myWins++;
    if (theirStrength > myStrength) theirWins++;
  }

  return myWins > theirWins ? 'We won' : theirWins > myWins ? 'We lost' : 'It was a tie';
}
```
