---
id: 6821ebda237de8297eaee792
title: "Завдання 16: перевернути круглі дужки"
challengeType: 28
dashedName: challenge-16
---

# --description--

Дано рядок, який містить правильно вкладені круглі дужки. Поверніть декодовану версію рядка, використовуючи такі правила:

- Усі символи всередині кожної пари круглих дужок потрібно перевернути.
- Круглі дужки слід видалити з фінального результату.
- Якщо дужки вкладені, спочатку потрібно перевернути найглибшу пару, а потім її результат включити у перевертання зовнішньої пари.
- Припустіть, що всі круглі дужки збалансовані та правильно вкладені.

# --hints--

`decode("(f(b(dc)e)a)")` має повертати `"abcdef"`.

```js
assert.equal(decode("(f(b(dc)e)a)"), "abcdef");
```

`decode("((is?)(a(t d)h)e(n y( uo)r)aC)")` має повертати `"Can you read this?"`.

```js
assert.equal(decode("((is?)(a(t d)h)e(n y( uo)r)aC)"), "Can you read this?");
```

`decode("f(Ce(re))o((e(aC)m)d)p")` має повертати `"freeCodeCamp"`.

```js
assert.equal(decode("f(Ce(re))o((e(aC)m)d)p"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(s) {

  return s;
}
```

# --solutions--

```js
function decode(s) {
  while (s.includes(')')) {
    const closeIndex = s.indexOf(')');
    const openIndex = s.lastIndexOf('(', closeIndex);
    const before = s.slice(0, openIndex);
    const group = s.slice(openIndex + 1, closeIndex).split('').reverse().join('');
    const after = s.slice(closeIndex + 1);
    s = before + group + after;
  }

  return s;
}
```
