---
id: 68216ef80f957572e7c340c5
title: "Завдання 12: декодер повідомлень"
challengeType: 28
dashedName: challenge-12
---

# --description--

Дано секретний рядок повідомлення та ціле число, яке представляє кількість літер, на яку було зсунуто повідомлення для його кодування. Поверніть розкодований рядок.

- Додатне число означає, що повідомлення було зсунуте вперед по алфавіту.
- Від’ємне число означає, що повідомлення було зсунуте назад по алфавіту.
- Регістр має значення, розкодовані символи повинні зберігати регістр своїх закодованих відповідників.
- Неалфавітні символи не слід розкодовувати.

# --hints--

`decode("Xlmw mw e wigvix qiwweki.", 4)` має повертати `"This is a secret message."`

```js
assert.equal(decode("Xlmw mw e wigvix qiwweki.", 4), "This is a secret message.");
```

`decode("Byffi Qilfx!", 20)` має повертати `"Hello World!"`

```js
assert.equal(decode("Byffi Qilfx!", 20), "Hello World!");
```

`decode("Zqd xnt njzx?", -1)` має повертати `"Are you okay?"`

```js
assert.equal(decode("Zqd xnt njzx?", -1), "Are you okay?");
```

`decode("oannLxmnLjvy", 9)` має повертати `"freeCodeCamp"`

```js
assert.equal(decode("oannLxmnLjvy", 9), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function decode(message, shift) {

  return message;
}
```

# --solutions--

```js
function decode(message, shift) {
  return message.split('').map(char => {
    if (/[a-zA-Z]/.test(char)) {
      const base = char === char.toLowerCase() ? 'a'.charCodeAt(0) : 'A'.charCodeAt(0);
      const charCode = char.charCodeAt(0);
      const offset = (charCode - base - shift + 26) % 26;
      return String.fromCharCode(base + offset);
    } else {
      return char;
    }
  }).join('');
}
```
