---
id: 6821ebfd237de8297eaee799
title: "Завдання 23: RGB у шістнадцятковий формат"
challengeType: 28
dashedName: challenge-23
---

# --description--

Дано рядок кольору CSS `rgb(r, g, b)`, поверніть його шістнадцятковий еквівалент.

Ось кілька прикладів вихідних даних для заданих вхідних:

| Вхідні дані | Вихідні дані |
|-------------|--------------|
| `"rgb(255, 255, 255)"`   | `"#ffffff"`    |
| `"rgb(1, 2, 3)"`   | `"#010203"`    |

- Зробіть усі літери у нижньому регістрі.
- Поверніть `#`, за яким іде шість символів. Не використовуйте скорочені значення.

# --hints--

`rgbToHex("rgb(255, 255, 255)")` має повертати `"#ffffff"`.

```js
assert.equal(rgbToHex("rgb(255, 255, 255)"), "#ffffff");
```

`rgbToHex("rgb(1, 11, 111)")` має повертати `"#010b6f"`.

```js
assert.equal(rgbToHex("rgb(1, 11, 111)"), "#010b6f");
```

`rgbToHex("rgb(173, 216, 230)")` має повертати `"#add8e6"`.

```js
assert.equal(rgbToHex("rgb(173, 216, 230)"), "#add8e6");
```

`rgbToHex("rgb(79, 123, 201)")` має повертати `"#4f7bc9"`.

```js
assert.equal(rgbToHex("rgb(79, 123, 201)"), "#4f7bc9");
```

# --seed--

## --seed-contents--

```js
function rgbToHex(rgb) {

  return rgb;
}
```

# --solutions--

```js
function rgbToHex(rgb) {
  const match = rgb.match(/\d+/g);
  const [r, g, b] = match.map(num =>
    Math.max(0, Math.min(255, parseInt(num)))
      .toString(16)
      .padStart(2, '0')
  );

  return `#${r}${g}${b}`;
}
```
