---
id: 6821ebe4237de8297eaee794
title: "Завдання 18: другий за вартістю"
challengeType: 28
dashedName: challenge-18
---

# --description--

Дано масив цілих чисел, які представляють ціну різних ноутбуків, та ціле число, що представляє ваш бюджет. Потрібно повернути:

1. Другий за вартістю ноутбук, якщо він у межах вашого бюджету, або
2. Найдорожчий ноутбук, який у межах вашого бюджету, або
3. `0`, якщо жоден ноутбук не відповідає вашому бюджету.

- Повторювані ціни слід ігнорувати.

# --hints--

`getLaptopCost([1500, 2000, 1800, 1400], 1900)` має повертати `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900)` має повертати `1800`

```js
assert.equal(getLaptopCost([1500, 2000, 2000, 1800, 1400], 1900), 1800);
```

`getLaptopCost([2099, 1599, 1899, 1499], 2200)` має повертати `1899`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 2200), 1899);
```

`getLaptopCost([2099, 1599, 1899, 1499], 1000)` має повертати `0`

```js
assert.equal(getLaptopCost([2099, 1599, 1899, 1499], 1000), 0);
```

`getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450)` має повертати `1400`

```js
assert.equal(getLaptopCost([1200, 1500, 1600, 1800, 1400, 2000], 1450), 1400);
```

# --seed--

## --seed-contents--

```js
function getLaptopCost(laptops, budget) {

  return laptops;
}
```

# --solutions--

```js
function getLaptopCost(laptops, budget) {
  laptops = [...new Set(laptops)].sort((a, b) => b - a);

  if (budget >= laptops[1]) return laptops[1];
  if (budget < laptops[laptops.length - 1]) return 0;

  for (let i = 2; i < laptops.length; i++) {
    if (budget >= laptops[i]) return laptops[i];
  }
}
```
