---
id: 6821ebd4237de8297eaee791
title: "Завдання 15: верблюжий регістр"
challengeType: 28
dashedName: challenge-15
---

# --description--

Дано рядок, поверніть його версію у верблюжому регістрі, використовуючи такі правила:

- Слова в рядковому аргументі розділені одним або кількома символами з наступної множини: пробіл (` `), дефіс (`-`) або підкреслення (`_`). Розглядайте будь-яку послідовність цих символів як розрив слова.
- Перше слово має бути повністю у нижньому регістрі.
- Кожне наступне слово має починатися з великої літери, а решта літер — у нижньому регістрі.
- Всі пробіли та роздільники потрібно видалити.

# --hints--

`toCamelCase("hello world")` має повернути `"helloWorld"`.

```js
assert.equal(toCamelCase("hello world"), "helloWorld");
```

`toCamelCase("HELLO WORLD")` має повернути `"helloWorld"`.

```js
assert.equal(toCamelCase("HELLO WORLD"), "helloWorld");
```

`toCamelCase("secret agent-X")` має повернути `"secretAgentX"`.

```js
assert.equal(toCamelCase("secret agent-X"), "secretAgentX");
```

`toCamelCase("FREE cODE cAMP")` має повернути `"freeCodeCamp"`.

```js
assert.equal(toCamelCase("FREE cODE cAMP"), "freeCodeCamp");
```

`toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk")` має повернути `"yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk"`.

```js
assert.equal(toCamelCase("ye old-_-sea  faring_buccaneer_-_with a - peg__leg----and a_parrot_ _named- _squawk"), "yeOldSeaFaringBuccaneerWithAPegLegAndAParrotNamedSquawk");
```

# --seed--

## --seed-contents--

```js
function toCamelCase(s) {

  return s;
}
```

# --solutions--

```js
function toCamelCase(s) {
  const words = s.replace(/[_\- ]+/g, ' ').split(' ');

  return words.map((word, i) => {
    if (i === 0) {
      return word.toLowerCase();
    } else {
      const tempWord = word.split('');
      return tempWord.shift().toUpperCase() + tempWord.join('').toLowerCase();
    }
  }).join('')
}
```
