---
id: 6821ebf3237de8297eaee797
title: "Завдання 21: генератор шістнадцяткових кольорів"
challengeType: 28
dashedName: challenge-21
---

# --description--

За заданим іменованим рядком кольору CSS згенеруйте випадковий шістнадцятковий (hex) код кольору, домінантний у заданому кольорі.

- Функція має приймати `"red"`, `"green"` або `"blue"` як вхідний аргумент.
- Якщо вхідні дані не є одним із цих значень, функція має повертати `"Invalid color"`.
- Функція має повертати випадковий шістизначний шістнадцятковий код кольору, де значення вхідного кольору є більшим за будь-яке інше.
- Приклад допустимих вихідних даних для заданого вхідного:

| Вхідні дані | Вихідні дані |
|-------------|--------------|
| `"red"`   | `"FF0000"`    |
| `"green"`   | `"00FF00"`    |
| `"blue"`   | `"0000FF"`    |

# --hints--

`generateHex("yellow")` має повертати `"Invalid color"`.

```js
assert.equal(generateHex("yellow"), "Invalid color");
```

`generateHex("red")` має повертати рядок довжиною шість символів.

```js
assert.lengthOf(generateHex("red"), 6);
```

`generateHex("red")` має повертати дійсний шістнадцятковий код кольору довжиною шість символів.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);
```

`generateHex("red")` має повертати дійсний шістнадцятковий код з більшим значенням червоного, ніж інші кольори.

```js
const hex = generateHex("red").toUpperCase();
const isValidHex = /^[0-9A-F]{6}$/.test(hex);
assert.isTrue(isValidHex);

const r = parseInt(hex.slice(0, 2), 16);
const g = parseInt(hex.slice(2, 4), 16);
const b = parseInt(hex.slice(4, 6), 16);

assert.isAbove(r, g);
assert.isAbove(r, b);
```

Виклик `generateHex("red")` двічі має повертати два різні шістнадцяткові значення кольору, де домінує червоний.

```js
const hex1 = generateHex("red").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(r1, g1);
assert.isAbove(r1, b1);

const hex2 = generateHex("red").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(r2, g2);
assert.isAbove(r2, b2);
assert.notEqual(hex1, hex2);
```

Виклик `generateHex("green")` двічі має повертати два різні шістнадцяткові значення кольору, де домінує зелений.

```js
const hex1 = generateHex("green").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(g1, r1);
assert.isAbove(g1, b1);

const hex2 = generateHex("green").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(g2, r2);
assert.isAbove(g2, b2);
assert.notEqual(hex1, hex2);
```

Виклик `generateHex("blue")` двічі має повертати два різні шістнадцяткові значення кольору, де домінує синій.

```js
const hex1 = generateHex("blue").toUpperCase();
const isValidHex1 = /^[0-9A-F]{6}$/.test(hex1)
assert.isTrue(isValidHex1);

const r1 = parseInt(hex1.slice(0, 2), 16);
const g1 = parseInt(hex1.slice(2, 4), 16);
const b1 = parseInt(hex1.slice(4, 6), 16);

assert.isAbove(b1, r1);
assert.isAbove(b1, g1);

const hex2 = generateHex("blue").toUpperCase();
const isValidHex2 = /^[0-9A-F]{6}$/.test(hex2);
assert.isTrue(isValidHex2);

const r2 = parseInt(hex2.slice(0, 2), 16);
const g2 = parseInt(hex2.slice(2, 4), 16);
const b2 = parseInt(hex2.slice(4, 6), 16);

assert.isAbove(b2, r2);
assert.isAbove(b2, g2);
assert.notEqual(hex1, hex2);
```

# --seed--

## --seed-contents--

```js
function generateHex(color) {

  return color;
}
```

# --solutions--

```js
function generateHex(color) {
  const toHex = n => n.toString(16).padStart(2, "0").toUpperCase();

  const dominant = Math.floor(Math.random() * 86) + 170;
  const weak1 = Math.floor(Math.random() * 170);
  const weak2 = Math.floor(Math.random() * 170);

  let r, g, b;

  switch (color) {
    case "red":
      r = dominant;
      g = weak1;
      b = weak2;
      break;
    case "green":
      r = weak1;
      g = dominant;
      b = weak2;
      break;
    case "blue":
      r = weak1;
      g = weak2;
      b = dominant;
      break;
    default:
      return "Invalid color";
  }

  return `${toHex(r)}${toHex(g)}${toHex(b)}`;
}
```
