---
id: 6821ebee237de8297eaee796
title: "Завдання 20: дублікати в масиві"
challengeType: 28
dashedName: challenge-20
---

# --description--

Дано масив цілих чисел. Поверніть масив цілих чисел, які зустрічаються більше одного разу в початковому масиві, відсортований за зростанням. Якщо жодне значення не зустрічається більше одного разу, поверніть порожній масив.

- У поверненому масиві має бути лише один екземпляр кожного значення.

# --hints--

`findDuplicates([1, 2, 3, 4, 5])` має повернути `[]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 5]), []);
```

`findDuplicates([1, 2, 3, 4, 1, 2])` має повернути `[1, 2]`.

```js
assert.deepEqual(findDuplicates([1, 2, 3, 4, 1, 2]), [1, 2]);
```

`findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4])` має повернути `[-6, 0, 2, 4, 5, 23]`.

```js
assert.deepEqual(findDuplicates([2, 34, 0, 1, -6, 23, 5, 3, 2, 5, 67, -6, 23, 2, 43, 2, 12, 0, 2, 4, 4]), [-6, 0, 2, 4, 5, 23]);
```

# --seed--

## --seed-contents--

```js
function findDuplicates(arr) {

  return arr;
}
```

# --solutions--

```js
function findDuplicates(arr) {
  const duplicates = [];
  
  for (let i = 0; i < arr.length - 1; i++) {
    for (let j = i + 1; j < arr.length; j++) {
      if (arr[i] === arr[j] && !duplicates.includes(arr[i])) {
        duplicates.push(arr[i]);
      }
    }
  }

  return duplicates.sort((a, b) => a - b);
}
```
