---
id: 6821ebea237de8297eaee795
title: "Завдання 19: Світло свічки"
challengeType: 28
dashedName: challenge-19
---

# --description--

Дано ціле число, яке представляє кількість свічок, з якими ви починаєте, та ціле число, яке показує, скільки використаних свічок потрібно, щоб створити нову. Поверніть кількість свічок, які ви використаєте після того, як створите і спалите якомога більше.

Наприклад, якщо дано 7 свічок і потрібно 2 використані свічки, щоб зробити нову:

1. Спаліть 7 свічок, щоб отримати 7 залишків,
2. Переробіть 6 залишків у 3 нові свічки (залишився 1 залишок),
3. Спаліть 3 свічки, щоб отримати ще 3 залишки (всього 4),
4. Переробіть 4 залишки у 2 нові свічки,
5. Спаліть 2 свічки, щоб отримати 2 залишки,
6. Переробіть 2 залишки у 1 нову свічку,
7. Спаліть 1 свічку.

У цьому прикладі ви спалите всього 13 свічок.

# --hints--

`burnCandles(7, 2)` має повертати `13`

```js
assert.equal(burnCandles(7, 2), 13);
```

`burnCandles(10, 5)` має повертати `12`

```js
assert.equal(burnCandles(10, 5), 12);
```

`burnCandles(20, 3)` має повертати `29`

```js
assert.equal(burnCandles(20, 3), 29);
```

`burnCandles(17, 4)` має повертати `22`

```js
assert.equal(burnCandles(17, 4), 22);
```

`burnCandles(2345, 3)` має повертати `3517`

```js
assert.equal(burnCandles(2345, 3), 3517);
```

# --seed--

## --seed-contents--

```js
function burnCandles(candles, leftoversNeeded) {

  return candles;
}
```

# --solutions--

```js
function burnCandles(candles, leftoversNeeded) {
  let totalBurned = 0;
  let unusedLeftovers = 0;

  while (candles > 0) {
    totalBurned += candles;
    const leftovers = candles + unusedLeftovers;
    candles = Math.floor(leftovers / leftoversNeeded);
    unusedLeftovers = leftovers % leftoversNeeded;
  }

  return totalBurned;
}
```
