---
id: 6821ec02237de8297eaee79a
title: "Завдання 24: Панграма"
challengeType: 28
dashedName: challenge-24
---

# --description--

Дано слово або речення та рядок із буквами в нижньому регістрі. Визначте, чи використовує слово або речення всі букви з заданої множини принаймні один раз і не містить інших букв.

- Ігноруйте неалфавітні символи у слові або реченні.
- Ігноруйте регістр букв у слові або реченні.

# --hints--

`isPangram("hello", "helo")` має повертати `true`

```js
assert.isTrue(isPangram("hello", "helo"));
```

`isPangram("hello", "hel")` має повертати `false`

```js
assert.isFalse(isPangram("hello", "hel"));
```

`isPangram("hello", "helow")` має повертати `false`

```js
assert.isFalse(isPangram("hello", "helow"));
```

`isPangram("hello world", "helowrd")` має повертати `true`

```js
assert.isTrue(isPangram("hello world", "helowrd"));
```

`isPangram("Hello World!", "helowrd")` має повертати `true`

```js
assert.isTrue(isPangram("Hello World!", "helowrd"));
```

`isPangram("Hello World!", "heliowrd")` має повертати `false`

```js
assert.isFalse(isPangram("Hello World!", "heliowrd"));
```

`isPangram("freeCodeCamp", "frcdmp")` має повертати `false`

```js
assert.isFalse(isPangram("freeCodeCamp", "frcdmp"));
```

`isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz")` має повертати `true`

```js
assert.isTrue(isPangram("The quick brown fox jumps over the lazy dog.", "abcdefghijklmnopqrstuvwxyz"));
```

# --seed--

## --seed-contents--

```js
function isPangram(sentence, letters) {

  return sentence;
}
```

# --solutions--

```js
function isPangram(sentence, letters) {
  const usedLetters = [];
  for (let i = 0; i < sentence.length; i++) {
    const letter = sentence[i].toLowerCase();
    if (!usedLetters.includes(letter) && /[a-z]/.test(letter)) {
      usedLetters.push(letter);
    }
  }

  const sortedLetters = letters.split('').sort().join('');
  const sortedUsedLetters = usedLetters.sort().join('');

  return sortedLetters === sortedUsedLetters;
}
```
