---
id: 6821ebc9237de8297eaee78f
title: "Завдання 13: неприродне просте число"
challengeType: 28
dashedName: challenge-13
---

# --description--

Дано ціле число, визначте, чи є це число простим числом або від’ємним простим числом.

- Просте число — це додатне ціле число більше за 1, яке ділиться лише на 1 та на саме себе.
- Від’ємне просте число — це від’ємна версія додатного простого числа.
- `1` та `0` не вважаються простими числами.

# --hints--

`isUnnaturalPrime(1)` має повертати `false`.

```js
assert.isFalse(isUnnaturalPrime(1));
```

`isUnnaturalPrime(-1)` має повертати `false`.

```js
assert.isFalse(isUnnaturalPrime(-1));
```

`isUnnaturalPrime(19)` має повертати `true`.

```js
assert.isTrue(isUnnaturalPrime(19));
```

`isUnnaturalPrime(-23)` має повертати `true`.

```js
assert.isTrue(isUnnaturalPrime(-23));
```

`isUnnaturalPrime(0)` має повертати `false`.

```js
assert.isFalse(isUnnaturalPrime(0));
```

`isUnnaturalPrime(97)` має повертати `true`.

```js
assert.isTrue(isUnnaturalPrime(97));
```

`isUnnaturalPrime(-61)` має повертати `true`.

```js
assert.isTrue(isUnnaturalPrime(-61));
```

`isUnnaturalPrime(99)` має повертати `false`.

```js
assert.isFalse(isUnnaturalPrime(99));
```

`isUnnaturalPrime(-44)` має повертати `false`.

```js
assert.isFalse(isUnnaturalPrime(-44));
```

# --seed--

## --seed-contents--

```js
function isUnnaturalPrime(n) {

  return n;
}
```

# --solutions--

```js
function isUnnaturalPrime(n) {
  const abs = Math.abs(n);

  if (abs <= 1) return false;

  for (let i = 2; i <= Math.sqrt(abs); i++) {
    if (abs % i === 0) return false;
  }

  return true;
}
```
