---
id: 6821ebf8237de8297eaee798
title: "Завдання 22: послідовність Трибоначчі"
challengeType: 28
dashedName: challenge-22
---

# --description--

Послідовність Трибоначчі — це ряд чисел, де кожне число є сумою трьох попередніх. Починаючи з `0`, `0` і `1`, перші 10 чисел послідовності — це `0`, `0`, `1`, `1`, `2`, `4`, `7`, `13`, `24`, `44`.

Дано масив, що містить перші три числа послідовності Трибоначчі, та ціле число, яке представляє довжину послідовності. Поверніть масив, що містить послідовність заданої довжини.

- Ваша функція має працювати з послідовностями будь-якої довжини, більшої або рівної нулю.
- Якщо довжина дорівнює нулю, поверніть порожній масив.
- Зверніть увагу, що початкові числа є частиною послідовності.

# --hints--

`tribonacciSequence([0, 0, 1], 20)` має повернути `[0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 20), [0, 0, 1, 1, 2, 4, 7, 13, 24, 44, 81, 149, 274, 504, 927, 1705, 3136, 5768, 10609, 19513]);
```

`tribonacciSequence([21, 32, 43], 1)` має повернути `[21]`.


```js
assert.deepEqual(tribonacciSequence([21, 32, 43], 1), [21]);
```

`tribonacciSequence([0, 0, 1], 0)` має повернути `[]`.

```js
assert.deepEqual(tribonacciSequence([0, 0, 1], 0), []);
```

`tribonacciSequence([10, 20, 30], 2)` має повернути `[10, 20]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 2), [10, 20]);
```

`tribonacciSequence([10, 20, 30], 3)` має повернути `[10, 20, 30]`.

```js
assert.deepEqual(tribonacciSequence([10, 20, 30], 3), [10, 20, 30]);
```

`tribonacciSequence([123, 456, 789], 8)` має повернути `[123, 456, 789, 1368, 2613, 4770, 8751, 16134]`.

```js
assert.deepEqual(tribonacciSequence([123, 456, 789], 8), [123, 456, 789, 1368, 2613, 4770, 8751, 16134]);
```

# --seed--

## --seed-contents--

```js
function tribonacciSequence(startSequence, length) {

  return length;
}
```

# --solutions--

```js
function tribonacciSequence(startSequence, length) {
  if (length === 0) return [];
  if (length === 1) return [startSequence[0]];
  if (length === 2) return [startSequence[0], startSequence[1]];
  if (length === 3) return [...startSequence];

  const sequence = [...startSequence];
  while (sequence.length < length) {
    const nextValue = sequence[sequence.length - 1] + sequence[sequence.length - 2] + + sequence[sequence.length - 3];
    sequence.push(nextValue);
  }
  return sequence;
}
```
