---
id: 68adce01c0e1144d0a90295e
title: "Завдання 29: побудова акроніма"
challengeType: 28
dashedName: challenge-29
---

# --description--

Дано рядок, що містить одне або кілька слів. Поверніть акронім зі слів, дотримуючись таких обмежень:

- Акронім повинен складатися з першої літери кожного слова, написаної великими літерами, якщо не вказано інше.
- Акронім має ігнорувати першу літеру цих слів, якщо вони не є першим словом у заданому рядку: `a`, `for`, `an`, `and`, `by` та `of`.
- Літери акроніма мають повертатися у тому порядку, в якому вони зустрічаються.
- Акронім не повинен містити пробілів.

# --hints--

`buildAcronym("Search Engine Optimization")` має повернути `"SEO"`.

```js
assert.equal(buildAcronym("Search Engine Optimization"), "SEO");
```

`buildAcronym("Frequently Asked Questions")` має повернути `"FAQ"`.

```js
assert.equal(buildAcronym("Frequently Asked Questions"), "FAQ");
```

`buildAcronym("National Aeronautics and Space Administration")` має повернути `"NASA"`.

```js
assert.equal(buildAcronym("National Aeronautics and Space Administration"), "NASA");
```

`buildAcronym("Federal Bureau of Investigation")` має повернути `"FBI"`.

```js
assert.equal(buildAcronym("Federal Bureau of Investigation"), "FBI");
```

`buildAcronym("For your information")` має повернути `"FYI"`.

```js
assert.equal(buildAcronym("For your information"), "FYI");
```

`buildAcronym("By the way")` має повернути `"BTW"`.

```js
assert.equal(buildAcronym("By the way"), "BTW");
```

`buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily")` має повернути `"AUHWPOTIMSH"`.

```js
assert.equal(buildAcronym("An unstoppable herd of waddling penguins overtakes the icy mountains and sings happily"), "AUHWPOTIMSH");
```

# --seed--

## --seed-contents--

```js
function buildAcronym(str) {

  return str;
}
```

# --solutions--

```js
function buildAcronym(str) {
  const smallWords = ["a", "for", "an", "and", "by", "of"];
  const words = str.split(" ");
  let acronym = "";

  for (let i = 0; i < words.length; i++) {
    const word = words[i];
    const lowerWord = word.toLowerCase();

    if (i === 0 || !smallWords.includes(lowerWord)) {
      acronym += word[0].toUpperCase();
    }
  }

  return acronym;
}
```
