---
id: 68adce01c0e1144d0a902958
title: "Завдання 26: перевірка IPv4"
challengeType: 28
dashedName: challenge-26
---

# --description--

Дано рядок, визначте, чи є він дійсною IPv4-адресою. Дійсна IPv4-адреса складається з чотирьох цілих чисел, розділених крапками (`.`). Кожне число має відповідати таким умовам:

- Воно знаходиться в діапазоні від 0 до 255 включно.
- Не має провідних нулів (наприклад, 0 дозволено, а 01 — ні).
- Дозволені лише числові символи.

# --hints--

`isValidIPv4("192.168.1.1")` має повернути `true`.

```js
assert.isTrue(isValidIPv4("192.168.1.1"));
```

`isValidIPv4("0.0.0.0")` має повернути `true`.

```js
assert.isTrue(isValidIPv4("0.0.0.0"));
```

`isValidIPv4("255.01.50.111")` має повернути `false`.

```js
assert.isFalse(isValidIPv4("255.01.50.111"));
```

`isValidIPv4("255.00.50.111")` має повернути `false`.

```js
assert.isFalse(isValidIPv4("255.00.50.111"));
```

`isValidIPv4("256.101.50.115")` має повернути `false`.

```js
assert.isFalse(isValidIPv4("256.101.50.115"));
```

`isValidIPv4("192.168.101.")` має повернути `false`.

```js
assert.isFalse(isValidIPv4("192.168.101."));
```

`isValidIPv4("192168145213")` має повернути `false`.

```js
assert.isFalse(isValidIPv4("192168145213"));
```

# --seed--

## --seed-contents--

```js
function isValidIPv4(ipv4) {

  return ipv4;
}
```

# --solutions--

```js
function isValidIPv4(ipv4) {
  const parts = ipv4.split(".");
  
  if (parts.length !== 4) return false;

  for (let part of parts) {
    if (!/^\d+$/.test(part)) return false;    

    const num = Number(part);
    if (num < 0 || num > 255) return false;

    if (part.length > 1 && part.startsWith("0")) return false;
  }

  return true;
}
```
