---
id: 68adce01c0e1144d0a902956
title: "Завдання 25: повторювач голосних"
challengeType: 28
dashedName: challenge-25
---

# --description--

Дано рядок, поверніть нову версію рядка, де кожна голосна повторюється на один раз більше, ніж попередня зустрінута голосна. Наприклад, перша голосна у реченні залишається без змін. Друга голосна з’являється двічі поспіль. Третя голосна з’являється тричі поспіль і так далі.

- Літери `a`, `e`, `i`, `o` та `u`, у верхньому або нижньому регістрі, вважаються голосними.
- Початкова голосна має зберігати свій регістр.
- Повторені голосні мають бути в нижньому регістрі.
- Всі символи, що не є голосними, мають зберігати свій початковий регістр.

# --hints--

`repeatVowels("hello world")` має повернути `"helloo wooorld"`.

```js
assert.equal(repeatVowels("hello world"), "helloo wooorld");
```

`repeatVowels("freeCodeCamp")` має повернути `"freeeCooodeeeeCaaaaamp"`.

```js
assert.equal(repeatVowels("freeCodeCamp"), "freeeCooodeeeeCaaaaamp");
```

`repeatVowels("AEIOU")` має повернути `"AEeIiiOoooUuuuu"`.

```js
assert.equal(repeatVowels("AEIOU"), "AEeIiiOoooUuuuu");
```

`repeatVowels("I like eating ice cream in Iceland")` має повернути `"I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand"`.

```js
assert.equal(repeatVowels("I like eating ice cream in Iceland"), "I liikeee eeeeaaaaatiiiiiing iiiiiiiceeeeeeee creeeeeeeeeaaaaaaaaaam iiiiiiiiiiin Iiiiiiiiiiiiceeeeeeeeeeeeelaaaaaaaaaaaaaand");
```

# --seed--

## --seed-contents--

```js
function repeatVowels(str) {

  return str;
}
```

# --solutions--

```js
function repeatVowels(str) {
  const vowels = "aeiouAEIOU";
  let count = 0;
  let result = "";

  for (let char of str) {
    result += char;
    if (vowels.includes(char)) {
      result += char.repeat(count).toLowerCase();
      count++;
    }
  }

  return result;
}
```
