---
id: 68adce01c0e1144d0a90295c
title: "Завдання 28: парсер римських чисел"
challengeType: 28
dashedName: challenge-28
---

# --description--

Дано рядок, що представляє римське число. Поверніть його ціле число.

Римські числа складаються з таких символів і значень:

| Символ | Значення |
|--------|----------|
| I      | 1        |
| V      | 5        |
| X      | 10       |
| L      | 50       |
| C      | 100      |
| D      | 500      |
| M      | 1000     |

- Числа читаються зліва направо. Якщо менше число стоїть перед більшим, його значення віднімається. Інакше значення додаються.

# --hints--

`parseRomanNumeral("III")` має повертати `3`.

```js
assert.equal(parseRomanNumeral("III"), 3);
```

`parseRomanNumeral("IV")` має повертати `4`.

```js
assert.equal(parseRomanNumeral("IV"), 4);
```

`parseRomanNumeral("XXVI")` має повертати `26`.

```js
assert.equal(parseRomanNumeral("XXVI"), 26);
```

`parseRomanNumeral("XCIX")` має повертати `99`.

```js
assert.equal(parseRomanNumeral("XCIX"), 99);
```

`parseRomanNumeral("CDLX")` має повертати `460`.

```js
assert.equal(parseRomanNumeral("CDLX"), 460);
```

`parseRomanNumeral("DIV")` має повертати `504`.

```js
assert.equal(parseRomanNumeral("DIV"), 504);
```

`parseRomanNumeral("MMXXV")` має повертати `2025`.

```js
assert.equal(parseRomanNumeral("MMXXV"), 2025);
```

# --seed--

## --seed-contents--

```js
function parseRomanNumeral(numeral) {

  return numeral;
}
```

# --solutions--

```js
function parseRomanNumeral(numeral) {
  const romanMap = {
    I: 1,
    V: 5,
    X: 10,
    L: 50,
    C: 100,
    D: 500,
    M: 1000
  };

  let total = 0;

  for (let i = 0; i < numeral.length; i++) {
    const current = romanMap[numeral[i]];
    const next = romanMap[numeral[i + 1]];

    if (next && current < next) {
      total -= current;
    } else {
      total += current;
    }
  }

  return total;
}
```
