---
id: 68af0687ef34c76c28ffa54b
title: "Завдання 32: Реверс речення"
challengeType: 28
dashedName: challenge-32
---

# --description--

Дано рядок слів, поверніть новий рядок зі словами у зворотному порядку. Наприклад, перше слово має бути в кінці повернутого рядка, а останнє слово — на початку повернутого рядка.

- У заданому рядку слова можуть бути розділені одним або кількома пробілами.
- Повернутий рядок повинен містити лише один пробіл між словами.

# --hints--

`reverseSentence("world hello")` має повернути `"hello world"`.

```js
assert.equal(reverseSentence("world hello"), "hello world");
```

`reverseSentence("push commit git")` має повернути `"git commit push"`.

```js
assert.equal(reverseSentence("push commit git"), "git commit push");
```

`reverseSentence("npm  install  sudo")` має повернути `"sudo install npm"`.

```js
assert.equal(reverseSentence("npm  install   apt    sudo"), "sudo apt install npm");
```

`reverseSentence("import    default   function  export")` має повернути `"export function default import"`.

```js
assert.equal(reverseSentence("import    default   function  export"), "export function default import");
```

# --seed--

## --seed-contents--

```js
function reverseSentence(sentence) {

  return sentence;
}
```

# --solutions--

```js
function reverseSentence(sentence) {
  return sentence.split(/\s+/).reverse().join(" ");
}
```
