---
id: 68af0687ef34c76c28ffa547
title: "Завдання 30: унікальні символи"
challengeType: 28
dashedName: challenge-30
---

# --description--

Дано рядок, визначте, чи всі символи в рядку унікальні.

- Літери верхнього та нижнього регістрів слід вважати різними символами.

# --hints--

`allUnique("abc")` має повертати `true`.

```js
assert.isTrue(allUnique("abc"));
```

`allUnique("aA")` має повертати `true`.

```js
assert.isTrue(allUnique("aA"));
```

`allUnique("QwErTy123!@")` має повертати `true`.

```js
assert.isTrue(allUnique("QwErTy123!@"));
```

`allUnique("~!@#$%^&*()_+")` має повертати `true`.

```js
assert.isTrue(allUnique("~!@#$%^&*()_+"));
```

`allUnique("hello")` має повертати `false`.

```js
assert.isFalse(allUnique("hello"));
```

`allUnique("freeCodeCamp")` має повертати `false`.

```js
assert.isFalse(allUnique("freeCodeCamp"));
```

`allUnique("!@#*$%^&*()aA")` має повертати `false`.

```js
assert.isFalse(allUnique("!@#*$%^&*()aA"));
```

# --seed--

## --seed-contents--

```js
function allUnique(str) {

  return str;
}
```

# --solutions--

```js
function allUnique(str) {
  let seen = '';
  for (let char of str) {
    if (seen.includes(char)) {
      return false;
    }
    seen += char;
  }
  return true;
}
```
