---
id: 68b06e589bf2273243814777
title: "Завдання 39: Заповнити бак"
challengeType: 28
dashedName: challenge-39
---

# --description--

Враховуючи розмір паливного бака, поточний рівень палива та ціну за галон, поверніть вартість повного заправлення бака.

- `tankSize` — це загальна місткість бака в галонах.
- `fuelLevel` — це поточна кількість палива в баку в галонах.
- `pricePerGallon` — це вартість одного галона палива.
- Повернене значення має бути округлене до двох десяткових знаків у форматі: `"$d.dd"`.

# --hints--

`costToFill(20, 0, 4.00)` має повертати `"$80.00"`.

```js
assert.equal(costToFill(20, 0, 4.00), "$80.00");
```

`costToFill(15, 10, 3.50)` має повертати `"$17.50"`.

```js
assert.equal(costToFill(15, 10, 3.50), "$17.50");
```

`costToFill(18, 9, 3.25)` має повертати `"$29.25"`.

```js
assert.equal(costToFill(18, 9, 3.25), "$29.25");
```

`costToFill(12, 12, 4.99)` має повертати `"$0.00"`.

```js
assert.equal(costToFill(12, 12, 4.99), "$0.00");
```

`costToFill(15, 9.5, 3.98)` має повертати `"$21.89"`.

```js
assert.equal(costToFill(15, 9.5, 3.98), "$21.89");
```

# --seed--

## --seed-contents--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {

  return tankSize;
}
```

# --solutions--

```js
function costToFill(tankSize, fuelLevel, pricePerGallon) {
  let gallonsNeeded = tankSize - fuelLevel;
  let cost = gallonsNeeded * pricePerGallon;
  return `$${cost.toFixed(2)}`;
}
```
