---
id: 68b06e589bf227324381476f
title: "Завдання 35: Частота слів"
challengeType: 28
dashedName: challenge-35
---

# --description--

Дано абзац, поверніть масив трьох найчастіше вживаних слів.

- Слова в абзаці будуть розділені пробілами.
- Ігноруйте регістр у заданому абзаці. Наприклад, вважайте `Hello` і `hello` одним і тим самим словом.
- Ігноруйте розділові знаки в заданому абзаці. Розділові знаки — це коми (`,`), крапки (`.`) та знаки оклику (`!`).
- Повернутий масив повинен містити всі слова в нижньому регістрі.
- Повернутий масив повинен бути відсортований за спаданням, починаючи з найчастіше вживаного слова.

# --hints--

`getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding")` має повернути `["coding", "python", "in"]`.

```js
assert.deepEqual(getWords("Coding in Python is fun because coding Python allows for coding in Python easily while coding"), ["coding", "python", "in"]);
```

`getWords("I like coding. I like testing. I love debugging!")` має повернути `["i", "like", "coding"]`.

```js
assert.deepEqual(getWords("I like coding. I like testing. I love debugging!"), ["i", "like", "coding"]);
```

`getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!")` має повернути `["debug", "test", "deploy"]`.

```js
assert.deepEqual(getWords("Debug, test, deploy. Debug, debug, test, deploy. Debug, test, test, deploy!"), ["debug", "test", "deploy"]);
```

# --seed--

## --seed-contents--

```js
function getWords(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function getWords(paragraph) {
  const cleaned = paragraph.replace(/[.,!]/g, "").toLowerCase();
  const words = cleaned.split(/\s+/);

  const freq = {};
  for (const word of words) {
    if (word) {
      freq[word] = (freq[word] || 0) + 1;
    }
  }

  const sorted = Object.keys(freq).sort((a, b) => freq[b] - freq[a]);
  return sorted.slice(0, 3);
}
```
