---
id: 68b06e589bf2273243814771
title: "Завдання 36: регулятор термостата"
challengeType: 28
dashedName: challenge-36
---

# --description--

Враховуючи поточну температуру в кімнаті та цільову температуру, поверніть рядок, який вказує, як регулювати температуру в кімнаті відповідно до цих умов:

- Поверніть `"heat"`, якщо поточна температура нижча за цільову.
- Поверніть `"cool"`, якщо поточна температура вища за цільову.
- Поверніть `"hold"`, якщо поточна температура дорівнює цільовій.

# --hints--

`adjustThermostat(68, 72)` має повертати `"heat"`.

```js
assert.equal(adjustThermostat(68, 72), "heat");
```

`adjustThermostat(75, 72)` має повертати `"cool"`.

```js
assert.equal(adjustThermostat(75, 72), "cool");
```

`adjustThermostat(72, 72)` має повертати `"hold"`.

```js
assert.equal(adjustThermostat(72, 72), "hold");
```

`adjustThermostat(-20.5, -10.1)` має повертати `"heat"`.

```js
assert.equal(adjustThermostat(-20.5, -10.1), "heat");
```

`adjustThermostat(100, 99.9)` має повертати `"cool"`.

```js
assert.equal(adjustThermostat(100, 99.9), "cool");
```

`adjustThermostat(0.0, 0.0)` має повертати `"hold"`.

```js
assert.equal(adjustThermostat(0.0, 0.0), "hold");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(temp, target) {

  return temp;
}
```

# --solutions--

```js
function adjustThermostat(temp, target) {
  if (temp < target) {
    return "heat";
  } else if (temp > target) {
    return "cool";
  } else {
    return "hold";
  }
}
```
