---
id: 68b06e589bf2273243814773
title: "Завдання 37: Капіталізація речень"
challengeType: 28
dashedName: challenge-37
---

# --description--

Дано абзац. Поверніть новий абзац, у якому перша літера кожного речення буде великою.

- Всі інші символи мають залишатися без змін.
- Речення можуть закінчуватися крапкою (`.`), одним або кількома знаками питання (`?`) або одним або кількома знаками оклику (`!`).

# --hints--

`capitalize("this is a simple sentence.")` має повертати `"This is a simple sentence."`.

```js
assert.equal(capitalize("this is a simple sentence."), "This is a simple sentence.");
```

`capitalize("hello world. how are you?")` має повертати `"Hello world. How are you?"`.

```js
assert.equal(capitalize("hello world. how are you?"), "Hello world. How are you?");
```

`capitalize("i did today's coding challenge... it was fun!!")` має повертати `"I did today's coding challenge... It was fun!!"`.

```js
assert.equal(capitalize("i did today's coding challenge... it was fun!!"), "I did today's coding challenge... It was fun!!");
```

`capitalize("crazy!!!strange???unconventional...sentences.")` має повертати `"Crazy!!!Strange???Unconventional...Sentences."`.

```js
assert.equal(capitalize("crazy!!!strange???unconventional...sentences."), "Crazy!!!Strange???Unconventional...Sentences.");
```

`capitalize("there's a space before this period . why is there a space before that period ?")` має повертати `"There's a space before this period . Why is there a space before that period ?"`.

```js
assert.equal(capitalize("there's a space before this period . why is there a space before that period ?"), "There's a space before this period . Why is there a space before that period ?");
```

# --seed--

## --seed-contents--

```js
function capitalize(paragraph) {

  return paragraph;
}
```

# --solutions--

```js
function capitalize(paragraph) {
  return paragraph.replace(/([.!?]+)(\s*)([a-z])/g, (match, punc, spaces, char) => {
    return punc + spaces + char.toUpperCase();
  })
  .replace(/^([a-z])/, (match, char) => char.toUpperCase());
}
```
