---
id: 68b06e589bf2273243814775
title: "Завдання 38: генератор slug"
challengeType: 28
dashedName: challenge-38
---

# --description--

Дано рядок. Поверніть версію цього рядка, дружню до URL, з урахуванням таких обмежень:

- Всі літери мають бути в нижньому регістрі.
- Всі символи, які не є літерами, числами або пробілами, потрібно видалити.
- Всі пробіли потрібно замінити на URL-кодований пробіл `%20`.
- Послідовні пробіли потрібно замінити на один `%20`.
- Повернутий рядок не повинен починатися або закінчуватися `%20`.

# --hints--

`generateSlug("helloWorld")` має повернути `"helloworld"`.

```js
assert.equal(generateSlug("helloWorld"), "helloworld");
```

`generateSlug("hello world!")` має повернути `"hello%20world"`.

```js
assert.equal(generateSlug("hello world!"), "hello%20world");
```

`generateSlug(" hello-world ")` має повернути `"helloworld"`.

```js
assert.equal(generateSlug(" hello-world "), "helloworld");
```

`generateSlug("hello  world")` має повернути `"hello%20world"`.

```js
assert.equal(generateSlug("hello  world"), "hello%20world");
```

`generateSlug("  ?H^3-1*1]0! W[0%R#1]D  ")` має повернути `"h3110%20w0r1d"`.

```js
assert.equal(generateSlug("  ?H^3-1*1]0! W[0%R#1]D  "), "h3110%20w0r1d");
```

# --seed--

## --seed-contents--

```js
function generateSlug(str) {

  return str;
}
```

# --solutions--

```js
function generateSlug(str) {
  let cleaned = str.replace(/[^a-zA-Z0-9 ]+/g, "");
  cleaned = cleaned.replace(/\s+/g, " ").trim();
  return cleaned.toLowerCase().replace(/ /g, "%20");
}
```
