---
id: 68b1f72371a5ac895ac70a06
title: "Завдання 42: зберігання відео"
challengeType: 28
dashedName: challenge-42
---

# --description--

За заданим розміром відео, одиницею виміру розміру відео, обсягом жорсткого диска та одиницею виміру жорсткого диска поверніть кількість відео, які жорсткий диск може зберегти, використовуючи такі обмеження:

- Одиниця виміру розміру відео може бути байтами (`"B"`), кілобайтами (`"KB"`), мегабайтами (`"MB"`) або гігабайтами (`"GB"`).
- Якщо не вказано одну з вищезазначених одиниць для відео, поверніть `"Invalid video unit"`.
- Одиниця виміру обсягу жорсткого диска може бути гігабайтами (`"GB"`) або терабайтами (`"TB"`).
- Якщо не вказано одну з вищезазначених одиниць для жорсткого диска, поверніть `"Invalid drive unit"`.
- Поверніть кількість цілих відео, які помістяться на диску.
- Використовуйте такі конверсії:

| Одиниця | Еквівалент |
|:-------:|:----------:|
| 1 B     |   1 B      |
| 1 KB    |   1000 B   |
| 1 MB    |   1000 KB  |
| 1 GB    |   1000 MB  |
| 1 TB    |   1000 GB  |

Наприклад, за аргументів `500`, `"MB"`, `100` та `"GB"` визначте, скільки відео по 500 MB поміститься на жорсткому диску обсягом 100 GB.

# --hints--

`numberOfVideos(500, "MB", 100, "GB")` має повернути `200`.

```js
assert.equal(numberOfVideos(500, "MB", 100, "GB"), 200);
```

`numberOfVideos(1, "TB", 10, "TB")` має повернути `"Invalid video unit"`.

```js
assert.equal(numberOfVideos(1, "TB", 10, "TB"), "Invalid video unit");
```

`numberOfVideos(2000, "MB", 100000, "MB")` має повернути `"Invalid drive unit"`.

```js
assert.equal(numberOfVideos(2000, "MB", 100000, "MB"), "Invalid drive unit");
```

`numberOfVideos(500000, "KB", 2, "TB")` має повернути `4000`.

```js
assert.equal(numberOfVideos(500000, "KB", 2, "TB"), 4000);
```

`numberOfVideos(1.5, "GB", 2.2, "TB")` має повернути `1466`.

```js
assert.equal(numberOfVideos(1.5, "GB", 2.2, "TB"), 1466);
```

# --seed--

## --seed-contents--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {

  return videoSize;
}
```

# --solutions--

```js
function numberOfVideos(videoSize, videoUnit, driveSize, driveUnit) {
  const videoUnits = { "KB": 1000, "MB": 1000 * 1000, "GB": 1000 * 1000 * 1000 };
  const driveUnits = { "GB": 1000 * 1000 * 1000, "TB": 1000 * 1000 * 1000 * 1000 };

  if (!videoUnits[videoUnit]) return "Invalid video unit";
  if (!driveUnits[driveUnit]) return "Invalid drive unit";

  const videoBytes = videoSize * videoUnits[videoUnit];
  const driveBytes = driveSize * driveUnits[driveUnit];

  return Math.floor(driveBytes / videoBytes);
}
```
