---
id: 68b1f72371a5ac895ac70a02
title: "Завдання 40: Зберігання фотографій"
challengeType: 28
dashedName: challenge-40
---

# --description--

Дано розмір фотографії в мегабайтах (MB) та обсяг жорсткого диска в гігабайтах (GB). Поверніть кількість фотографій, які жорсткий диск може зберегти, враховуючи такі умови:

- 1 гігабайт дорівнює 1000 мегабайтам.
- Поверніть кількість цілих фотографій, які може зберегти диск.

# --hints--

`numberOfPhotos(1, 1)` має повертати `1000`.

```js
assert.equal(numberOfPhotos(1, 1), 1000);
```

`numberOfPhotos(2, 1)` має повертати `500`.

```js
assert.equal(numberOfPhotos(2, 1), 500);
```

`numberOfPhotos(4, 256)` має повертати `64000`.

```js
assert.equal(numberOfPhotos(4, 256), 64000);
```

`numberOfPhotos(3.5, 750)` має повертати `214285`.

```js
assert.equal(numberOfPhotos(3.5, 750), 214285);
```

`numberOfPhotos(3.5, 5.5)` має повертати `1571`.

```js
assert.equal(numberOfPhotos(3.5, 5.5), 1571);
```

# --seed--

## --seed-contents--

```js
function numberOfPhotos(photoSizeMb, hardDriveSizeGb) {

  return photoSizeMb;
}
```

# --solutions--

```js
function numberOfPhotos(photoSizeMb, driveSizeGb) {
  const driveSizeMb = driveSizeGb * 1000;
  return Math.floor(driveSizeMb / photoSizeMb);
}
```
