---
id: 68b1f72371a5ac895ac70a04
title: "Завдання 41: зберігання файлів"
challengeType: 28
dashedName: challenge-41
---

# --description--

За заданим розміром файлу, одиницею виміру розміру файлу та обсягом жорсткого диска в гігабайтах (GB) поверніть кількість файлів, які жорсткий диск може зберегти, використовуючи такі обмеження:

- Одиниця виміру розміру файлу може бути байтами (`"B"`), кілобайтами (`"KB"`) або мегабайтами (`"MB"`).
- Поверніть кількість цілих файлів, які помістяться на диску.
- Використовуйте такі конверсії:

| Одиниця | Еквівалент |
|:-------:|:----------:|
| 1 B     |   1 B      |
| 1 KB    |   1000 B   |
| 1 MB    |   1000 KB  |
| 1 GB    |   1000 MB  |

Наприклад, за аргументів `500`, `"KB"` і `1` визначте, скільки файлів розміром 500 KB поміститься на жорсткому диску обсягом 1 GB.

# --hints--

`numberOfFiles(500, "KB", 1)` має повернути `2000`.

```js
assert.equal(numberOfFiles(500, "KB", 1), 2000);
```

`numberOfFiles(50000, "B", 1)` має повернути `20000`.

```js
assert.equal(numberOfFiles(50000, "B", 1), 20000);
```

`numberOfFiles(5, "MB", 1)` має повернути `200`.

```js
assert.equal(numberOfFiles(5, "MB", 1), 200);
```

`numberOfFiles(4096, "B", 1.5)` має повернути `366210`.

```js
assert.equal(numberOfFiles(4096, "B", 1.5), 366210);
```

`numberOfFiles(220.5, "KB", 100)` має повернути `453514`.

```js
assert.equal(numberOfFiles(220.5, "KB", 100), 453514);
```

`numberOfFiles(4.5, "MB", 750)` має повернути `166666`.

```js
assert.equal(numberOfFiles(4.5, "MB", 750), 166666);
```

# --seed--

## --seed-contents--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {

  return fileSize;
}
```

# --solutions--

```js
function numberOfFiles(fileSize, fileUnit, driveSizeGb) {
  const driveSizeBytes = driveSizeGb * 1000 * 1000 * 1000;

  let fileSizeBytes;
  if (fileUnit === "B") {
    fileSizeBytes = fileSize;
  } else if (fileUnit === "KB") {
    fileSizeBytes = fileSize * 1000;
  } else {
    fileSizeBytes = fileSize * 1000 * 1000;
  }

  return Math.floor(driveSizeBytes / fileSizeBytes);
}
```
