---
id: 68b1f72371a5ac895ac70a0a
title: "Завдання 44: дзеркальний рядок"
challengeType: 28
dashedName: challenge-44
---

# --description--

Дано два рядки, визначте, чи є другий рядок дзеркальним відображенням першого.

- Рядок вважається дзеркальним, якщо він містить ті ж літери у зворотному порядку.
- Враховуйте верхній і нижній регістр літер як різні.
- Ігноруйте всі символи, що не є літерами алфавіту.

# --hints--

`isMirror("helloworld", "helloworld")` має повертати `false`.

```js
assert.isFalse(isMirror("helloworld", "helloworld"));
```

`isMirror("Hello World", "dlroW olleH")` має повертати `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW olleH"));
```

`isMirror("RaceCar", "raCecaR")` має повертати `true`.

```js
assert.isTrue(isMirror("RaceCar", "raCecaR"));
```

`isMirror("RaceCar", "RaceCar")` має повертати `false`.

```js
assert.isFalse(isMirror("RaceCar", "RaceCar"));
```

`isMirror("Mirror", "rorrim")` має повертати `false`.

```js
assert.isFalse(isMirror("Mirror", "rorrim"));
```

`isMirror("Hello World", "dlroW-olleH")` має повертати `true`.

```js
assert.isTrue(isMirror("Hello World", "dlroW-olleH"));
```

`isMirror("Hello World", "!dlroW !olleH")` має повертати `true`.

```js
assert.isTrue(isMirror("Hello World", "!dlroW !olleH"));
```

# --seed--

## --seed-contents--

```js
function isMirror(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function isMirror(str1, str2) {
  const clean1 = str1.replace(/[^a-zA-Z]/g, "");
  const clean2 = str2.replace(/[^a-zA-Z]/g, "");
  return clean1.split("").reverse().join("") === clean2;
}
```
