---
id: 68b1f72371a5ac895ac70a08
title: "Завдання 43: цифри проти літер"
challengeType: 28
dashedName: challenge-43
---

# --description--

Дано рядок. Поверніть `"digits"`, якщо в рядку більше цифр, ніж літер, `"letters"`, якщо більше літер, ніж цифр, і `"tie"`, якщо кількість цифр і літер однакова.

- Цифри складаються з `0-9`.
- Літери складаються з `a-z` у верхньому або нижньому регістрі.
- Ігноруйте всі інші символи.

# --hints--

`digitsOrLetters("abc123")` має повертати `"tie"`.

```js
assert.equal(digitsOrLetters("abc123"), "tie");
```

`digitsOrLetters("a1b2c3d")` має повертати `"letters"`.

```js
assert.equal(digitsOrLetters("a1b2c3d"), "letters");
```

`digitsOrLetters("1a2b3c4")` має повертати `"digits"`.

```js
assert.equal(digitsOrLetters("1a2b3c4"), "digits");
```

`digitsOrLetters("abc123!@#DEF")` має повертати `"letters"`.

```js
assert.equal(digitsOrLetters("abc123!@#DEF"), "letters");
```

`digitsOrLetters("H3110 W0R1D")` має повертати `"digits"`.

```js
assert.equal(digitsOrLetters("H3110 W0R1D"), "digits");
```

`digitsOrLetters("P455W0RD")` має повертати `"tie"`.

```js
assert.equal(digitsOrLetters("P455W0RD"), "tie");
```

# --seed--

## --seed-contents--

```js
function digitsOrLetters(str) {

  return str;
}
```

# --solutions--

```js
function digitsOrLetters(str) {
  let digitCount = 0;
  let letterCount = 0;

  for (let char of str) {
    if (/[0-9]/.test(char)) digitCount++;
    else if (/[a-zA-Z]/.test(char)) letterCount++;
  }

  if (digitCount > letterCount) return "digits";
  if (letterCount > digitCount) return "letters";
  return "tie";
}
```
