---
id: 68b7cadffed0e75a517da673
title: "Завдання 52: з двійкової у десяткову"
challengeType: 28
dashedName: challenge-52
---

# --description--

Дано рядок, що представляє двійкове число. Поверніть його десятковий еквівалент у вигляді числа.

Двійкове число використовує лише цифри `0` і `1` для представлення будь-якого числа. Щоб конвертувати двійкове число у десяткове, помножте кожну цифру на ступінь `2` і додайте їх разом. Починайте з множення правої крайньої цифри на `2^0`, наступної зліва — на `2^1` і так далі. Коли всі цифри будуть помножені на ступінь `2`, додайте отримані значення.

Наприклад, двійкове число `101` дорівнює `5` у десятковій системі, тому що:

```mathml
1 * 2^2 + 0 * 2^1 + 1 * 2^0 = 4 + 0 + 1 = 5
```

# --hints--

`toDecimal("101")` має повернути `5`.

```js
assert.equal(toDecimal("101"), 5);
```

`toDecimal("1010")` має повернути `10`.

```js
assert.equal(toDecimal("1010"), 10);
```

`toDecimal("10010")` має повернути `18`.

```js
assert.equal(toDecimal("10010"), 18);
```

`toDecimal("1010101")` має повернути `85`.

```js
assert.equal(toDecimal("1010101"), 85);
```

# --seed--

## --seed-contents--

```js
function toDecimal(binary) {

  return binary;
}
```

# --solutions--

```js
function toDecimal(binary) {
  let decimal = 0;
  for (let i = 0; i < binary.length; i++) {
    const bit = parseInt(binary[binary.length - 1 - i], 10);
    decimal += bit * Math.pow(2, i);
  }
  return decimal;
}
```
