---
id: 68b7cadffed0e75a517da675
title: "Завдання 53: десяткове число у двійкове"
challengeType: 28
dashedName: challenge-53
---

# --description--

Дано невід’ємне ціле число, поверніть його двійкове представлення у вигляді рядка.

Двійкове число використовує лише цифри `0` та `1` для представлення будь-якого числа. Щоб конвертувати десяткове число у двійкове, послідовно діліть число на `2` і записуйте остачу. Повторюйте, доки число не стане нулем. Читайте остачі у зворотному порядку — від останньої записаної до першої. Наприклад, щоб конвертувати `12` у двійкове:

```mathml
12 ÷ 2 = 6 remainder 0
6 ÷ 2 = 3 remainder 0
3 ÷ 2 = 1 remainder 1
1 ÷ 2 = 0 remainder 1
```

`12` у двійковому вигляді — це `1100`.

# --hints--

`toBinary(5)` має повернути `"101"`.

```js
assert.equal(toBinary(5), "101");
```

`toBinary(12)` має повернути `"1100"`.

```js
assert.equal(toBinary(12), "1100");
```

`toBinary(50)` має повернути `"110010"`.

```js
assert.equal(toBinary(50), "110010");
```

`toBinary(99)` має повернути `"1100011"`.

```js
assert.equal(toBinary(99), "1100011");
```

# --seed--

## --seed-contents--

```js
function toBinary(decimal) {

  return decimal;
}
```

# --solutions--

```js
function toBinary(decimal) {
  if (decimal === 0) return "0";
  let binary = "";
  while (decimal > 0) {
    binary = (decimal % 2) + binary;
    decimal = Math.floor(decimal / 2);
  }
  return binary;
}
```
