---
id: 68b7cadffed0e75a517da677
title: "Завдання 54: сила пароля P@ssw0rd Str3ngth!"
challengeType: 28
dashedName: challenge-54
---

# --description--

Дано рядок пароля, поверніть `"weak"`, `"medium"` або `"strong"` залежно від сили пароля.

Пароль оцінюється за такими правилами:

- Він має бути довжиною не менше 8 символів.
- Він містить як літери верхнього, так і нижнього регістру.
- Він містить принаймні одне число.
- Він містить принаймні один спеціальний символ із цієї множини: `!`, `@`, `#`, `$`, `%`, `^`, `&` або `*`.

Поверніть `"weak"`, якщо пароль відповідає менше ніж двом правилам.
Поверніть `"medium"`, якщо пароль відповідає 2 або 3 правилам.
Поверніть `"strong"`, якщо пароль відповідає всім 4 правилам.

# --hints--

`checkStrength("123456")` має повертати `"weak"`.

```js
assert.equal(checkStrength("123456"), "weak");
```

`checkStrength("pass!!!")` має повертати `"weak"`.

```js
assert.equal(checkStrength("pass!!!"), "weak");
```

`checkStrength("Qwerty")` має повертати `"weak"`.

```js
assert.equal(checkStrength("Qwerty"), "weak");
```

`checkStrength("PASSWORD")` має повертати `"weak"`.

```js
assert.equal(checkStrength("PASSWORD"), "weak");
```

`checkStrength("PASSWORD!")` має повертати `"medium"`.

```js
assert.equal(checkStrength("PASSWORD!"), "medium");
```

`checkStrength("PassWord%^!")` має повертати `"medium"`.

```js
assert.equal(checkStrength("PassWord%^!"), "medium");
```

`checkStrength("qwerty12345")` має повертати `"medium"`.

```js
assert.equal(checkStrength("qwerty12345"), "medium");
```

`checkStrength("S3cur3P@ssw0rd")` має повертати `"strong"`.

```js
assert.equal(checkStrength("S3cur3P@ssw0rd"), "strong");
```

`checkStrength("C0d3&Fun!")` має повертати `"strong"`.

```js
assert.equal(checkStrength("C0d3&Fun!"), "strong");
```

# --seed--

## --seed-contents--

```js
function checkStrength(password) {

  return password;
}
```

# --solutions--

```js
function checkStrength(password) {
  let rulesMet = 0;

  if (password.length >= 8) rulesMet++;
  if (/[a-z]/.test(password) && /[A-Z]/.test(password)) rulesMet++;
  if (/\d/.test(password)) rulesMet++;
  if (/[!@#$%^&*]/.test(password)) rulesMet++;

  if (rulesMet < 2) return "weak";
  if (rulesMet < 4) return "medium";
  return "strong";
}
```
