---
id: 68b7687dded630607aceccab
title: "Завдання 45: ідеальний квадрат"
challengeType: 28
dashedName: challenge-45
---

# --description--

Дано ціле число, визначте, чи є воно ідеальним квадратом.

- Число є ідеальним квадратом, якщо можна помножити ціле число саме на себе, щоб отримати це число. Наприклад, 9 — це ідеальний квадрат, бо можна помножити 3 на 3, щоб отримати 9.

# --hints--

`isPerfectSquare(9)` має повертати `true`.

```js
assert.isTrue(isPerfectSquare(9));
```

`isPerfectSquare(49)` має повертати `true`.

```js
assert.isTrue(isPerfectSquare(49));
```

`isPerfectSquare(1)` має повертати `true`.

```js
assert.isTrue(isPerfectSquare(1));
```

`isPerfectSquare(2)` має повертати `false`.

```js
assert.isFalse(isPerfectSquare(2));
```

`isPerfectSquare(99)` має повертати `false`.

```js
assert.isFalse(isPerfectSquare(99));
```

`isPerfectSquare(-9)` має повертати `false`.

```js
assert.isFalse(isPerfectSquare(-9));
```

`isPerfectSquare(0)` має повертати `true`.

```js
assert.isTrue(isPerfectSquare(0));
```

`isPerfectSquare(25281)` має повертати `true`.

```js
assert.isTrue(isPerfectSquare(25281));
```

# --seed--

## --seed-contents--

```js
function isPerfectSquare(n) {

  return n;
}
```

# --solutions--

```js
function isPerfectSquare(n) {
  if (n < 0) return false;
  const root = Math.floor(Math.sqrt(n));
  return root * root === n;
}
```
