---
id: 68b7687dded630607aceccb3
title: "Завдання 49: парсер заголовків CSV"
challengeType: 28
dashedName: challenge-49
---

# --description--

Дано першу лінію CSV-файлу (значення, розділені комами). Поверніть масив, що містить заголовки.

- Перша лінія CSV-файлу містить заголовки, розділені комами.
- Видаліть усі провідні або кінцеві пробіли з кожного заголовка.

# --hints--

`getHeadings("name,age,city")` має повернути `["name", "age", "city"]`.

```js
assert.deepEqual(getHeadings("name,age,city"), ["name", "age", "city"]);
```

`getHeadings("first name,last name,phone")` має повернути `["first name", "last name", "phone"]`.

```js
assert.deepEqual(getHeadings("first name,last name,phone"), ["first name", "last name", "phone"]);
```

`getHeadings("username , email , signup date ")` має повернути `["username", "email", "signup date"]`.

```js
assert.deepEqual(getHeadings("username , email , signup date "), ["username", "email", "signup date"]);
```

# --seed--

## --seed-contents--

```js
function getHeadings(csv) {

  return csv;
}
```

# --solutions--

```js
function getHeadings(csv) {
  return csv
    .split(",")
    .map(h => h.trim());
}
```
