---
id: 68b7687dded630607aceccaf
title: "Завдання 47: Спіймані на перевищенні швидкості"
challengeType: 28
dashedName: challenge-47
---

# --description--

Дано масив чисел, які представляють швидкість, з якою спостерігалися транспортні засоби, та число, що представляє обмеження швидкості. Поверніть масив із двох елементів: кількість транспортних засобів, які перевищували швидкість, а також середнє значення перевищення швидкості серед цих транспортних засобів.

- Якщо транспортних засобів, які перевищували швидкість, не було, поверніть `[0, 0]`.

# --hints--

`speeding([50, 60, 55], 60)` має повертати `[0, 0]`.

```js
assert.deepEqual(speeding([50, 60, 55], 60), [0, 0]);
```

`speeding([58, 50, 60, 55], 55)` має повертати `[2, 4]`.

```js
assert.deepEqual(speeding([58, 50, 60, 55], 55), [2, 4]);
```

`speeding([61, 81, 74, 88, 65, 71, 68], 70)` має повертати `[4, 8.5]`.

```js
assert.deepEqual(speeding([61, 81, 74, 88, 65, 71, 68], 70), [4, 8.5]);
```

`speeding([100, 105, 95, 102], 100)` має повертати `[2, 3.5]`.

```js
assert.deepEqual(speeding([100, 105, 95, 102], 100), [2, 3.5]);
```

`speeding([40, 45, 44, 50, 112, 39], 55)` має повертати `[1, 57]`.

```js
assert.deepEqual(speeding([40, 45, 44, 50, 112, 39], 55), [1, 57]);
```

# --seed--

## --seed-contents--

```js
function speeding(speeds, limit) {

  return speeds;
}
```

# --solutions--

```js
function speeding(speeds, limit) {
  let speeding = 0;
  let totalOver = 0;

  for (let speed of speeds) {
    if (speed > limit) {
      speeding++;
      totalOver += (speed - limit);
    }
  }

  if (speeding === 0) return [0, 0];

  return [speeding, totalOver / speeding];
}
```
