---
id: 68b7687dded630607aceccb1
title: "Завдання 48: детектор спаму"
challengeType: 28
dashedName: challenge-48
---

# --description--

Дано номер телефону у форматі `"+A (BBB) CCC-DDDD"`, де кожна літера представляє цифру за таким правилом:

- `A` представляє код країни і може містити будь-яку кількість цифр.
- `BBB` представляє код області і завжди містить три цифри.
- `CCC` та `DDDD` представляють місцевий номер і завжди містять відповідно три та чотири цифри.

Визначте, чи є цей номер спамом за такими критеріями:

- Код країни довший за 2 цифри або не починається з нуля (`0`).
- Код області більший за 900 або менший за 200.
- Сума перших трьох цифр місцевого номера зустрічається серед останніх чотирьох цифр місцевого номера.
- У номері одна й та сама цифра повторюється чотири або більше разів поспіль (ігноруючи символи форматування).

# --hints--

`isSpam("+0 (200) 234-0182")` має повернути `false`.

```js
assert.isFalse(isSpam("+0 (200) 234-0182"));
```

`isSpam("+091 (555) 309-1922")` має повернути `true`.

```js
assert.isTrue(isSpam("+091 (555) 309-1922"));
```

`isSpam("+1 (555) 435-4792")` має повернути `true`.

```js
assert.isTrue(isSpam("+1 (555) 435-4792"));
```

`isSpam("+0 (955) 234-4364")` має повернути `true`.

```js
assert.isTrue(isSpam("+0 (955) 234-4364"));
```

`isSpam("+0 (155) 131-6943")` має повернути `true`.

```js
assert.isTrue(isSpam("+0 (155) 131-6943"));
```

`isSpam("+0 (555) 135-0192")` має повернути `true`.

```js
assert.isTrue(isSpam("+0 (555) 135-0192"));
```

`isSpam("+0 (555) 564-1987")` має повернути `true`.

```js
assert.isTrue(isSpam("+0 (555) 564-1987"));
```

`isSpam("+00 (555) 234-0182")` має повернути `false`.

```js
assert.isFalse(isSpam("+00 (555) 234-0182"));
```

# --seed--

## --seed-contents--

```js
function isSpam(number) {

  return number;
}
```

# --solutions--

```js
function isSpam(number) {
  const match = number.match(/^\+(\d+)\s\((\d{3})\)\s(\d{3})-(\d{4})$/);
  const [, countryCode, areaCode, ccc, dddd] = match;
  const allDigits = countryCode + areaCode + ccc + dddd;
  
  if (countryCode.length > 2 || !countryCode.startsWith("0")) return true;

  const areaNum = parseInt(areaCode, 10);
  if (areaNum > 900 || areaNum < 200) return true;

  const sumCCC = ccc.split("").reduce((a, b) => a + parseInt(b), 0);
  if (dddd.includes(sumCCC.toString())) return true;

  if (/(\d)\1\1\1/.test(allDigits)) return true;

  return false;
}
```
