---
id: 68c1a929005bf54d342aa8d2
title: "Завдання 55: тиждень космосу, день 1: класифікація зірок"
challengeType: 28
dashedName: challenge-55
---

# --description--

4 жовтня починається Всесвітній тиждень космосу. Наступні сім днів принесуть вам завдання з програмування на астрономічну тематику.

Для сьогоднішнього завдання вам дають температуру поверхні зірки в Кельвінах (K) і потрібно визначити її зоряну класифікацію на основі таких діапазонів:

- `"O"`: 30 000 K або вище
- `"B"`: від 10 000 K до 29 999 K
- `"A"`: від 7 500 K до 9 999 K
- `"F"`: від 6 000 K до 7 499 K
- `"G"`: від 5 200 K до 5 999 K
- `"K"`: від 3 700 K до 5 199 K
- `"M"`: від 0 K до 3 699 K

- Поверніть класифікацію заданої зірки.

# --hints--

`classification(5778)` має повертати `"G"`.

```js
assert.equal(classification(5778), "G");
```

`classification(2400)` має повертати `"M"`.

```js
assert.equal(classification(2400), "M");
```

`classification(9999)` має повертати `"A"`.

```js
assert.equal(classification(9999), "A");
```

`classification(3700)` має повертати `"K"`.

```js
assert.equal(classification(3700), "K");
```

`classification(3699)` має повертати `"M"`.

```js
assert.equal(classification(3699), "M");
```

`classification(210000)` має повертати `"O"`.

```js
assert.equal(classification(210000), "O");
```

`classification(6000)` має повертати `"F"`.

```js
assert.equal(classification(6000), "F");
```

`classification(11432)` має повертати `"B"`.

```js
assert.equal(classification(11432), "B");
```

# --seed--

## --seed-contents--

```js
function classification(temp) {

  return temp;
}
```

# --solutions--

```js
function classification(temp) {
  if (temp >= 30000) return "O";
  if (temp >= 10000) return "B";
  if (temp >= 7500) return "A";
  if (temp >= 6000) return "F";
  if (temp >= 5200) return "G";
  if (temp >= 3700) return "K";
  return "M";
}
```
