---
id: 68c1a929005bf54d342aa8d4
title: "Завдання 57: тиждень космосу, день 3: телефон додому"
challengeType: 28
dashedName: challenge-57
---

# --description--

На третій день тижня космосу вам надано масив чисел, які представляють відстані (у кілометрах) між вами, супутниками та вашою рідною планетою у маршруті зв’язку. Визначте, скільки часу займе повідомлення, надіслане через цей маршрут, щоб досягти планети призначення, враховуючи такі умови:

- Перше значення в масиві — це відстань від вашого місця розташування до першого супутника.
- Кожне наступне значення, крім останнього, — це відстань до наступного супутника.
- Останнє значення в масиві — це відстань від попереднього супутника до вашої рідної планети.
- Повідомлення рухається зі швидкістю 300 000 км/с.
- Кожен супутник, через який проходить повідомлення, додає затримку передачі в 0,5 секунди.
- Поверніть число, округлене до 4 десяткових знаків, без зайвих нулів у кінці.

# --hints--

`sendMessage([300000, 300000])` має повертати `2.5`.

```js
assert.equal(sendMessage([300000, 300000]), 2.5);
```

`sendMessage([384400, 384400])` має повертати `3.0627`.

```js
assert.equal(sendMessage([384400, 384400]), 3.0627);
```

`sendMessage([54600000, 54600000])` має повертати `364.5`.

```js
assert.equal(sendMessage([54600000, 54600000]), 364.5);
```

`sendMessage([1000000, 500000000, 1000000])` має повертати `1674.3333`.

```js
assert.equal(sendMessage([1000000, 500000000, 1000000]), 1674.3333);
```

`sendMessage([10000, 21339, 50000, 31243, 10000])` має повертати `2.4086`.

```js
assert.equal(sendMessage([10000, 21339, 50000, 31243, 10000]), 2.4086);
```

`sendMessage([802101, 725994, 112808, 3625770, 481239])` має повертати `21.1597`.

```js
assert.equal(sendMessage([802101, 725994, 112808, 3625770, 481239]), 21.1597);
```

# --seed--

## --seed-contents--

```js
function sendMessage(route) {

  return route;
}
```

# --solutions--

```js
function sendMessage(route) {
  let totalDistance = route.reduce((a, d) => a += d, 0);

  const delay = (route.length - 1) * 0.5
  const time = totalDistance / 300000
  const total = time + delay;
  return Number(total.toFixed(4))
}
```
