---
id: 68c1a929005bf54d342aa8d6
title: "Завдання 59: Тиждень космосу, день 5: зона Голділок"
challengeType: 28
dashedName: challenge-59
---

# --description--

На п’ятий день Тижня космосу ви обчислите «зону Голділок» зірки — область навколо зірки, де умови «саме такі», щоб існувала рідка вода.

За заданою масою зірки поверніть масив із початковою та кінцевою відстанями її зони Голділок у астрономічних одиницях.

Щоб обчислити зону Голділок:

1. Знайдіть світність зірки, піднявши її масу у степінь 3.5.
2. Початок зони — це 0.95 помножене на квадратний корінь зі світності.
3. Кінець зони — це 1.37 помножене на квадратний корінь зі світності.

- Поверніть відстані, округлені до двох десяткових знаків.

Наприклад, за маси `1` поверніть `[0.95, 1.37]`.

# --hints--

`goldilocksZone(1)` має повернути `[0.95, 1.37]`.

```js
assert.deepEqual(goldilocksZone(1), [0.95, 1.37]);
```

`goldilocksZone(0.5)` має повернути `[0.28, 0.41]`.

```js
assert.deepEqual(goldilocksZone(0.5), [0.28, 0.41]);
```

`goldilocksZone(6)` має повернути `[21.85, 31.51]`.

```js
assert.deepEqual(goldilocksZone(6), [21.85, 31.51]);
```

`goldilocksZone(3.7)` має повернути `[9.38, 13.52]`.

```js
assert.deepEqual(goldilocksZone(3.7), [9.38, 13.52]);
```

`goldilocksZone(20)` має повернути `[179.69, 259.13]`.

```js
assert.deepEqual(goldilocksZone(20), [179.69, 259.13]);
```

# --seed--

## --seed-contents--

```js
function goldilocksZone(mass) {

  return mass;
}
```

# --solutions--

```js
function goldilocksZone(mass) {
  const luminosity = Math.pow(mass, 3.5);
  const start = 0.95 * Math.sqrt(luminosity);
  const end = 1.37 * Math.sqrt(luminosity);
  return [Number(start.toFixed(2)), Number(end.toFixed(2))];
}
```
