---
id: 68c1a929005bf54d342aa8d5
title: "Завдання 58: тиждень космосу, день 4: місце посадки"
challengeType: 28
dashedName: challenge-58
---

# --description--

У четвертий день тижня космосу вам дають матрицю чисел (масив масивів), що представляє потенційні місця посадки для вашого ровера. Знайдіть найбезпечніше місце посадки за такими правилами:

- Кожна клітинка матриці містить число від `0-9`, включно.
- Будь-яке `0` представляє потенційне місце посадки.
- Будь-яке число, окрім `0`, занадто небезпечне для посадки. Чим більше число, тим небезпечніше.
- Найбезпечніше місце визначається як клітинка `0`, навколо якої (до 4 сусідів, діагоналі ігноруємо) сума небезпеки найменша.
- Ігноруйте сусідів за межами матриці (кути та краї мають менше сусідів).
- Поверніть індекси найбезпечнішого місця посадки. Завжди буде лише одне найбезпечніше місце.

Наприклад, для такого вхідного значення:

```js
[
  [1, 0],
  [2, 0]
]
```

Поверніть `[0, 1]`, індекси `0` у першому масиві.

# --hints--

`findLandingSpot([[1, 0], [2, 0]])` має повернути `[0, 1]`.

```js
assert.deepEqual(findLandingSpot([[1, 0], [2, 0]]), [0, 1]);
```

`findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]])` має повернути `[1, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 0, 3], [7, 0, 4], [8, 0, 5]]), [1, 1]);
```

`findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]])` має повернути `[2, 2]`.

```js
assert.deepEqual(findLandingSpot([[1, 2, 1], [0, 0, 2], [3, 0, 0]]), [2, 2]);
```

`findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]])` має повернути `[2, 1]`.

```js
assert.deepEqual(findLandingSpot([[9, 6, 0, 8], [7, 1, 1, 0], [3, 0, 3, 9], [8, 6, 0, 9]]), [2, 1]);
```

# --seed--

## --seed-contents--

```js
function findLandingSpot(matrix) {

  return matrix;
}
```

# --solutions--

```js
function findLandingSpot(matrix) {
  let bestSpot = null;
  let lowestNeighborSum = Infinity;

  for (let i = 0; i < matrix.length; i++) {
    for (let j = 0; j < matrix[i].length; j++) {

      if (matrix[i][j] === 0) {
        let currentNeighborSum = 0;

        if (i > 0) currentNeighborSum += matrix[i - 1][j];
        if (j < matrix[i].length - 1) currentNeighborSum += matrix[i][j + 1];
        if (i < matrix.length - 1) currentNeighborSum += matrix[i + 1][j];
        if (j > 0) currentNeighborSum += matrix[i][j - 1];

        if (currentNeighborSum < lowestNeighborSum) {
          lowestNeighborSum = currentNeighborSum;
          bestSpot = [i, j];
        }
      }
    }
  }

  return bestSpot;
}
```
