---
id: 68c1a929005bf54d342aa8d3
title: "Завдання 56: тиждень космосу, день 2: пошук екзопланет"
challengeType: 28
dashedName: challenge-56
---

# --description--

Для другого дня тижня космосу вам дається рядок, де кожен символ представляє показник яскравості зірки. Визначте, чи виявили ці показники екзопланету за допомогою методу транзиту. Метод транзиту полягає в тому, що планета проходить перед зіркою, зменшуючи її спостережувану яскравість.

- Показники яскравості складаються лише з символів `0-9` і `A-Z`, де кожен показник відповідає таким числовим значенням:
- Символи `0-9` відповідають рівням яскравості `0-9`.
- Символи `A-Z` відповідають рівням яскравості `10-35`.

Зірка вважається такою, що має екзопланету, якщо будь-який окремий показник менший або дорівнює 80% від середнього значення всіх показників. Наприклад, якщо середня яскравість зірки дорівнює 10, її вважатимуть такою, що має екзопланету, якщо будь-який окремий показник 8 або менше.

# --hints--

`hasExoplanet("665544554")` має повертати `false`.

```js
assert.isFalse(hasExoplanet("665544554"));
```

`hasExoplanet("FGFFCFFGG")` має повертати `true`.

```js
assert.isTrue(hasExoplanet("FGFFCFFGG"));
```

`hasExoplanet("MONOPLONOMONPLNOMPNOMP")` має повертати `false`.

```js
assert.isFalse(hasExoplanet("MONOPLONOMONPLNOMPNOMP"));
```

`hasExoplanet("FREECODECAMP")` має повертати `true`.

```js
assert.isTrue(hasExoplanet("FREECODECAMP"));
```

`hasExoplanet("9AB98AB9BC98A")` має повертати `false`.

```js
assert.isFalse(hasExoplanet("9AB98AB9BC98A"));
```

`hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE")` має повертати `true`.

```js
assert.isTrue(hasExoplanet("ZXXWYZXYWYXZEGZXWYZXYGEE"));
```

# --seed--

## --seed-contents--

```js
function hasExoplanet(readings) {

  return readings;
}
```

# --solutions--

```js
function hasExoplanet(readings) {
  let total = 0;
  const values = readings.split('').map(c => parseInt(c, 36));
  
  total = values.reduce((sum, v) => sum + v, 0);
  const average = total / values.length;
  const threshold = average * 0.8;

  return values.some(v => v <= threshold);
}
```
