---
id: 68c497f3aaefc9fd9f1b0e24
title: "Завдання 60: Тиждень космосу, день 6: фаза місяця"
challengeType: 28
dashedName: challenge-60
---

# --description--

На шостий день Тижня космосу вам буде надано дату у форматі `"YYYY-MM-DD"`, і потрібно визначити фазу місяця на цей день за такими правилами:

Використовуйте спрощений місячний цикл тривалістю 28 днів, поділений на чотири рівні фази:

- `"New"`: дні з 1 по 7
- `"Waxing"`: дні з 8 по 14
- `"Full"`: дні з 15 по 21
- `"Waning"`: дні з 22 по 28

Після 28-го дня цикл повторюється з 1-го дня — новим місяцем.

- Використовуйте `"2000-01-06"` як референсний новий місяць (1-й день циклу) для визначення фази заданого дня.
- Вам не будуть надані дати раніше за референсну.
- Поверніть правильну фазу у вигляді рядка.

**Примітка:** 1-й день означає день нового місяця, тобто минуло 0 днів від останнього нового місяця.

# --hints--

`moonPhase("2000-01-12")` має повернути `"New"`.

```js
assert.equal(moonPhase("2000-01-12"), "New");
```

`moonPhase("2000-01-13")` має повернути `"Waxing"`.

```js
assert.equal(moonPhase("2000-01-13"), "Waxing");
```

`moonPhase("2014-10-15")` має повернути `"Full"`.

```js
assert.equal(moonPhase("2014-10-15"), "Full");
```

`moonPhase("2012-10-21")` має повернути `"Waning"`.

```js
assert.equal(moonPhase("2012-10-21"), "Waning");
```

`moonPhase("2022-12-14")` має повернути `"New"`.

```js
assert.equal(moonPhase("2022-12-14"), "New");
```

# --seed--

## --seed-contents--

```js
function moonPhase(dateString) {

  return dateString;
}
```

# --solutions--

```js
function moonPhase(dateString) {
  const ONE_DAY_IN_MS = 1000 * 60 * 60 * 24;
  const refDate = new Date("2000-01-06");
  const targetDate = new Date(dateString);
  const diffMs = targetDate - refDate;
  const diffDays = diffMs / ONE_DAY_IN_MS;
  const phaseDay = (diffDays % 28 + 1);

  if (phaseDay <= 7) {
    return "New";
  } else if (phaseDay <= 14) {
    return "Waxing";
  } else if (phaseDay <= 21) {
    return "Full"
  } else {
    return "Waning";
  }
}

```
