---
id: 68c497f3aaefc9fd9f1b0e26
title: "Завдання 62: з шістнадцяткової у десяткову"
challengeType: 28
dashedName: challenge-62
---

# --description--

Дано рядок, що представляє шістнадцяткове число (основа 16). Поверніть його десяткове (основа 10) значення як ціле число.

Шістнадцяткова система числення використовує 16 цифр:

- `0-9` представляють значення від `0` до `9`.
- `A-F` представляють значення від `10` до `15`.

Ось часткова таблиця перетворень:

| Шістнадцяткове | Десяткове |
|:--------------:|:---------:|
| 0              | 0         |
| 1              | 1         |
| ...            | ...       |
| 9              | 9         |
| A              | 10        |
| ...            | ...       |
| F              | 15        |
| 10             | 16        |
| ...            | ...       |
| 9F             | 159       |
| A0             | 160       |
| ...            | ...       |
| FF             | 255       |
| 100            | 256       |

- Рядок міститиме лише символи `0–9` та `A–F`.

# --hints--

`hexToDecimal("A")` має повернути `10`.

```js
assert.equal(hexToDecimal("A"), 10);
```

`hexToDecimal("15")` має повернути `21`.

```js
assert.equal(hexToDecimal("15"), 21);
```

`hexToDecimal("2E")` має повернути `46`.

```js
assert.equal(hexToDecimal("2E"), 46);
```

`hexToDecimal("FF")` має повернути `255`.

```js
assert.equal(hexToDecimal("FF"), 255);
```

`hexToDecimal("A3F")` має повернути `2623`.

```js
assert.equal(hexToDecimal("A3F"), 2623);
```

# --seed--

## --seed-contents--

```js
function hexToDecimal(hex) {

  return hex;
}
```

# --solutions--

```js
function hexToDecimal(hex) {

  return parseInt(hex, 16);
}
```
