---
id: 68cae5b538ff798bbd4da002
title: "Завдання 64: 24-годинний формат у 12-годинний"
challengeType: 28
dashedName: challenge-64
---

# --description--

Дано рядок, який представляє час доби у 24-годинному форматі `"HHMM"`. Поверніть час у його еквівалентному 12-годинному форматі `"H:MM AM"` або `"H:MM PM"`.

- Вхідні дані завжди будуть рядком з чотирьох цифр у 24-годинному форматі, від `"0000"` до `"2359"`.

# --hints--

`to12("1124")` має повернути `"11:24 AM"`.

```js
assert.equal(to12("1124"), "11:24 AM");
```

`to12("0900")` має повернути `"9:00 AM"`.

```js
assert.equal(to12("0900"), "9:00 AM");
```

`to12("1455")` має повернути `"2:55 PM"`.

```js
assert.equal(to12("1455"), "2:55 PM");
```

`to12("2346")` має повернути `"11:46 PM"`.

```js
assert.equal(to12("2346"), "11:46 PM");
```

`to12("0030")` має повернути `"12:30 AM"`.

```js
assert.equal(to12("0030"), "12:30 AM");
```

# --seed--

## --seed-contents--

```js
function to12(time) {

  return time;
}
```

# --solutions--

```js
function convertHours(hours) {
  if (hours === 0) return 12;
  if (hours > 12) return hours - 12;
  return hours;
}

function to12(time) {
  const hours = parseInt(time.slice(0, 2), 10);
  const minutes = time.slice(2);
  const period = hours < 12 ? 'AM' : 'PM';

  return `${convertHours(hours)}:${minutes} ${period}`;
}
```
