---
id: 68cae5b538ff798bbd4da005
title: "Завдання 67: перевірка е-пошти"
challengeType: 28
dashedName: challenge-67
---

# --description--

Дано рядок, визначте, чи є він дійсною адресою е-пошти за такими обмеженнями:

- Він повинен містити рівно один символ `@`.
- Локальна частина (до `@`):
  - Може містити лише літери (`a-z`, `A-Z`), цифри (`0-9`), крапки (`.`), підкреслення (`_`) або дефіси (`-`).
  - Не може починатися або закінчуватися крапкою.
- Доменна частина (після `@`):
  - Має містити принаймні одну крапку.
  - Має закінчуватися крапкою, за якою йдуть щонайменше дві літери.
- Жодна з частин — локальна чи доменна — не може містити дві крапки поспіль.

# --hints--

`validate("a@b.cd")` має повернути `true`.

```js
assert.isTrue(validate("a@b.cd"));
```

`validate("hell.-w.rld@example.com")` має повернути `true`.

```js
assert.isTrue(validate("hell.-w.rld@example.com"));
```

`validate(".b@sh.rc")` має повернути `false`.

```js
assert.isFalse(validate(".b@sh.rc"));
```

`validate("example@test.c0")` має повернути `false`.

```js
assert.isFalse(validate("example@test.c0"));
```

`validate("freecodecamp.org")` має повернути `false`.

```js
assert.isFalse(validate("freecodecamp.org"));
```

`validate("develop.ment_user@c0D!NG.R.CKS")` має повернути `true`.

```js
assert.isTrue(validate("develop.ment_user@c0D!NG.R.CKS"));
```

`validate("hello.@wo.rld")` має повернути `false`.

```js
assert.isFalse(validate("hello.@wo.rld"));
```

`validate("hello@world..com")` має повернути `false`.

```js
assert.isFalse(validate("hello@world..com"));
```

`validate("develop..ment_user@c0D!NG.R.CKS")` має повернути `false`.

```js
assert.isFalse(validate("develop..ment_user@c0D!NG.R.CKS"));
```

`validate("git@commit@push.io")` має повернути `false`.

```js
assert.isFalse(validate("git@commit@push.io"));
```

# --seed--

## --seed-contents--

```js
function validate(email) {

  return email;
}
```

# --solutions--

```js
function validate(email) {
  if (email.includes('..')) return false;

  const parts = email.split('@');
  if (parts.length !== 2) return false;

  const [local, domain] = parts;

  if (local.startsWith('.') || local.endsWith('.')) return false;
  if (!/^[a-zA-Z0-9._-]+$/.test(local)) return false;
  if (!domain.includes('.')) return false;

  const tld = domain.split('.').pop();
  if (tld.length < 2 || !/^[a-zA-Z]+$/.test(tld)) return false;

  return true;
}
```
