---
id: 68cae5b538ff798bbd4da007
title: "Завдання 69: Відсутні шкарпетки"
challengeType: 28
dashedName: challenge-69
---

# --description--

Дано ціле число, яке представляє кількість пар шкарпеток, з якими ви почали, та інше ціле число, яке представляє, скільки циклів прання ви пройшли. Поверніть кількість повних пар шкарпеток, які у вас зараз є, використовуючи такі умови:

- Кожні 2 цикли прання ви втрачаєте одну шкарпетку.
- Кожні 3 цикли прання ви знаходите одну відсутню шкарпетку.
- Кожні 5 циклів прання одна шкарпетка зношується і її потрібно викинути.
- Кожні 10 циклів прання ви купуєте пару шкарпеток.
- Загальна кількість шкарпеток ніколи не може бути меншою за нуль.
- Правила можуть накладатися. Наприклад, у 10-му циклі прання ви втратите одну шкарпетку, викинете одну шкарпетку і купите нову пару шкарпеток.
- Поверніть кількість повних пар шкарпеток.

# --hints--

`sockPairs(2, 5)` має повертати `1`.

```js
assert.equal(sockPairs(2, 5), 1);
```

`sockPairs(1, 2)` має повертати `0`.

```js
assert.equal(sockPairs(1, 2), 0);
```

`sockPairs(5, 11)` має повертати `4`.

```js
assert.equal(sockPairs(5, 11), 4);
```

`sockPairs(6, 25)` має повертати `3`.

```js
assert.equal(sockPairs(6, 25), 3);
```

`sockPairs(1, 8)` має повертати `0`.

```js
assert.equal(sockPairs(1, 8), 0);
```

# --seed--

## --seed-contents--

```js
function sockPairs(pairs, cycles) {

  return pairs;
}
```

# --solutions--

```js
function sockPairs(pairs, cycles) {
  let socks = pairs * 2;

  for (let i = 1; i <= cycles; i++) {
    if (i % 2 === 0) socks -= 1;
    if (i % 3 === 0) socks += 1;
    if (i % 5 === 0) socks -= 1;
    if (i % 10 === 0) socks += 2;

    if (socks < 0) socks = 0;
  }

  return Math.floor(socks / 2);
}
```
