---
id: 68cae5b538ff798bbd4da001
title: "Завдання 63: битва слів"
challengeType: 28
dashedName: challenge-63
---

# --description--

Дано два речення, які представляють вашу команду та команду суперника, де кожне слово вашої команди змагається з відповідним словом команди суперника. Визначте, яка команда перемагає, використовуючи такі правила:

- У заданих реченнях завжди буде однакова кількість слів.
- Слова розділені одним пробілом і містять лише літери.
- Значення кожного слова — це сума значень його літер.
- Літери `a` до `z` відповідають значенням `1` до `26`. Наприклад, `a` має значення `1`, а `z` — `26`.
- Велика літера подвоює значення літери. Наприклад, `A` має значення `2`, а `Z` — `52`.
- Слова змагаються послідовно: перше слово вашої команди змагається з першим словом команди суперника, і так далі.
- Слово перемагає, якщо його значення більше за значення слова суперника.
- Команда з більшою кількістю виграшних слів є переможцем.

Поверніть `"We win"`, якщо ваша команда перемагає, `"We lose"`, якщо ваша команда програє, і `"Draw"`, якщо обидві команди мають однакову кількість перемог.

# --hints--

`battle("hello world", "hello word")` має повертати `"We win"`.

```js
assert.equal(battle("hello world", "hello word"), "We win");
```

`battle("Hello world", "hello world")` має повертати `"We win"`.

```js
assert.equal(battle("Hello world", "hello world"), "We win");
```

`battle("lorem ipsum", "kitty ipsum")` має повертати `"We lose"`.

```js
assert.equal(battle("lorem ipsum", "kitty ipsum"), "We lose");
```

`battle("hello world", "world hello")` має повертати `"Draw"`.

```js
assert.equal(battle("hello world", "world hello"), "Draw");
```

`battle("git checkout", "git switch")` має повертати `"We win"`.

```js
assert.equal(battle("git checkout", "git switch"), "We win");
```

`battle("Cheeseburger with fries", "Cheeseburger with Fries")` має повертати `"We lose"`.

```js
assert.equal(battle("Cheeseburger with fries", "Cheeseburger with Fries"), "We lose");
```

`battle("We must never surrender", "Our team must win")` має повертати `"Draw"`.

```js
assert.equal(battle("We must never surrender", "Our team must win"), "Draw");
```

# --seed--

## --seed-contents--

```js
function battle(ourTeam, opponent) {

  return ourTeam;
}
```

# --solutions--

```js
function getCharacterValue(char) {
  const lowerCaseValue = char.toLowerCase().charCodeAt(0) - 96;
  return char >= 'A' && char <= 'Z' ? lowerCaseValue * 2 : lowerCaseValue;
}

function getWordValue(word) {
  let wordValue = 0;
  for (let i=0; i<word.length; i++) {
    wordValue += getCharacterValue(word[i])
  }

  return wordValue;
}

function battle(ourTeam, opponent) {
  const myWords = ourTeam.split(' ');
  const opponentWords = opponent.split(' ');

  let myWins = 0, opponentWins = 0;

  for (let i=0; i<myWords.length; i++) {
    const myWordValue = getWordValue(myWords[i]);
    const opponentWordValue = getWordValue(opponentWords[i])

    if (myWordValue > opponentWordValue) myWins++;
    if (opponentWordValue > myWordValue) opponentWins++;
  }

  if (myWins > opponentWins) return 'We win';
  if (opponentWins > myWins) return 'We lose';
  return 'Draw';
}
```
