---
id: 68cae5b538ff798bbd4da003
title: "Завдання 65: підрахунок рядків"
challengeType: 28
dashedName: challenge-65
---

# --description--

Дано два рядки, визначте, скільки разів другий рядок зустрічається в першому.

- Шаблон рядка може перекриватися в першому рядку. Наприклад, `"aaa"` містить `"aa"` двічі. Перші два `a` і другі два.

# --hints--

`count('abcdefg', 'def')` має повернути `1`.

```js
assert.equal(count('abcdefg', 'def'), 1);
```

`count('hello', 'world')` має повернути `0`.

```js
assert.equal(count('hello', 'world'), 0);
```

`count('mississippi', 'iss')` має повернути `2`.

```js
assert.equal(count('mississippi', 'iss'), 2);
```

`count('she sells seashells by the seashore', 'sh')` має повернути `3`.

```js
assert.equal(count('she sells seashells by the seashore', 'sh'), 3);
```

`count('101010101010101010101', '101')` має повернути `10`.

```js
assert.equal(count('101010101010101010101', '101'), 10);
```

# --seed--

## --seed-contents--

```js
function count(text, pattern) {

  return text;
}
```

# --solutions--

```js
function count(text, pattern) {
  let occurrences = 0;

  for (let i = 0; i <= text.length - pattern.length; i++) {
    if (text.slice(i, i + pattern.length) === pattern) {
      occurrences++;
    }
  }

  return occurrences;
}
```
