---
id: 68cae5b538ff798bbd4da008
title: "Завдання 70: витягування атрибутів HTML"
challengeType: 28
dashedName: challenge-70
---

# --description--

Дано рядок із дійсним HTML-елементом, поверніть атрибути цього елемента за такими критеріями:

- Вам буде надано лише один елемент.
- Атрибути будуть у форматі: `attribute="value"`.
- Поверніть масив рядків, де кожен рядок містить властивість атрибута та його значення, розділені комою, у такому форматі: `["attribute1, value1", "attribute2, value2"]`.
- Поверніть атрибути у тому порядку, в якому вони задані.
- Якщо атрибути не знайдені, поверніть порожній масив.

# --hints--

`extractAttributes('<span class="red"></span>')` має повертати `["class, red"]`.

```js
assert.deepEqual(extractAttributes('<span class="red"></span>'), ["class, red"]);
```

`extractAttributes('<meta charset="UTF-8" />')` має повертати `["charset, UTF-8"]`.

```js
assert.deepEqual(extractAttributes('<meta charset="UTF-8" />'), ["charset, UTF-8"]);
```

`extractAttributes("<p>Lorem ipsum dolor sit amet</p>")` має повертати `[]`.

```js
assert.deepEqual(extractAttributes("<p>Lorem ipsum dolor sit amet</p>"), []);
```

`extractAttributes('<input name="email" type="email" required="true" />')` має повертати `["name, email", "type, email", "required, true"]`.

```js
assert.deepEqual(extractAttributes('<input name="email" type="email" required="true" />'), ["name, email", "type, email", "required, true"]);
```

`extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>')` має повертати `["id, submit", "class, btn btn-primary"]`.

```js
assert.deepEqual(extractAttributes('<button id="submit" class="btn btn-primary">Submit</button>'), ["id, submit", "class, btn btn-primary"]);
```

# --seed--

## --seed-contents--

```js
function extractAttributes(element) {

  return element;
}
```

# --solutions--

```js
function extractAttributes(element) {
  const regex = /(\w[\w-]*)="([^"]*)"/g;
  const result = [];
  let match;

  while ((match = regex.exec(element)) !== null) {
    const attr = match[1];
    const val = match[2];
    result.push(`${attr}, ${val}`);
  }

  return result;
}
```
