---
id: 68cae5b538ff798bbd4da004
title: "Завдання 66: видалення HTML-тегів"
challengeType: 28
dashedName: challenge-66
---

# --description--

Дано рядок HTML-коду, видаліть теги та поверніть звичайний текстовий вміст.

- Вхідний рядок міститиме лише дійсний HTML.
- HTML-теги можуть бути вкладеними.
- Видаліть теги та будь-які атрибути.

Наприклад, `'<a href="#">Click here</a>'` має повернути `"Click here"`.

# --hints--

`stripTags('<a href="#">Click here</a>')` має повернути `"Click here"`.

```js
assert.equal(stripTags('<a href="#">Click here</a>'), "Click here");
```

`stripTags('<p class="center">Hello <b>World</b>!</p>')` має повернути `"Hello World!"`.

```js
assert.equal(stripTags('<p class="center">Hello <b>World</b>!</p>'), "Hello World!");
```

`stripTags('<img src="cat.jpg" alt="Cat">')` має повернути порожній рядок (`""`).

```js
assert.equal(stripTags('<img src="cat.jpg" alt="Cat">'), "");
```

`stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>')` має повернути `sectionsection`.

```js
assert.equal(stripTags('<main id="main"><section class="section">section</section><section class="section">section</section></main>'), "sectionsection");
```

# --seed--

## --seed-contents--

```js
function stripTags(html) {

  return html;
}
```

# --solutions--

```js
function stripTags(html) {
  return html.replace(/<[^>]*>/g, '');
}
```
