---
id: 68cae5b538ff798bbd4da009
title: "Завдання 71: калькулятор чайових"
challengeType: 28
dashedName: challenge-71
---

# --description--

Враховуючи ціну вашої страви та кастомний відсоток чайових, поверніть масив із трьома значеннями чайових: 15%, 20% і кастомною сумою.

- Ціни будуть задані у форматі: `"$N.NN"`.
- Кастомні відсотки чайових будуть задані у такому форматі: `"25%"`.
- Поверніть суми у тому ж форматі `"$N.NN"`, округлені до двох десяткових знаків.

Наприклад, якщо ціна страви `"$10.00"`, а кастомне значення чайових `"25%"`, поверніть `["$1.50", "$2.00", "$2.50"]`.

# --hints--

`calculateTips("$10.00", "25%")` має повернути `["$1.50", "$2.00", "$2.50"]`.

```js
assert.deepEqual(calculateTips("$10.00", "25%"), ["$1.50", "$2.00", "$2.50"]);
```

`calculateTips("$89.67", "26%")` має повернути `["$13.45", "$17.93", "$23.31"]`.

```js
assert.deepEqual(calculateTips("$89.67", "26%"), ["$13.45", "$17.93", "$23.31"]);
```

`calculateTips("$19.85", "9%")` має повернути `["$2.98", "$3.97", "$1.79"]`.

```js
assert.deepEqual(calculateTips("$19.85", "9%"), ["$2.98", "$3.97", "$1.79"]);
```

# --seed--

## --seed-contents--

```js
function calculateTips(mealPrice, customTip) {

  return mealPrice;
}
```

# --solutions--

```js
function calculateTips(mealPrice, customTip) {
  const meal = parseFloat(mealPrice.slice(1));
  const customPercent = parseFloat(customTip.slice(0, -1));
  const tips = [15, 20, customPercent];

  return tips.map(percent => {
    const amount = (meal * percent / 100).toFixed(2);
    return `$${amount}`;
  });
}
```
