---
id: 68cae5b538ff798bbd4da00a
title: "Завдання 72: регулятор термостата 2"
challengeType: 28
dashedName: challenge-72
---

# --description--

Враховуючи поточну температуру в кімнаті у Фаренгейтах та цільову температуру у Цельсіях, поверніть рядок, який вказує, як регулювати температуру в кімнаті відповідно до цих умов:

- Поверніть `"Heat: X degrees Fahrenheit"`, якщо поточна температура нижча за цільову. `X` — це кількість градусів у Фаренгейтах, на яку потрібно нагріти кімнату, щоб досягти цілі, округлена до 1 десяткового знака.
- Поверніть `"Cool: X degrees Fahrenheit"`, якщо поточна температура вища за цільову. `X` — це кількість градусів у Фаренгейтах, на яку потрібно охолодити кімнату, щоб досягти цілі, округлена до 1 десяткового знака.
- Поверніть `"Hold"`, якщо поточна температура дорівнює цільовій.

Щоб конвертувати температуру з Цельсія у Фаренгейти, помножте температуру в Цельсіях на 1.8 і додайте 32 до результату (`F = (C * 1.8) + 32`).

# --hints--

`adjustThermostat(32, 0)` має повертати `"Hold"`.

```js
assert.equal(adjustThermostat(32, 0), "Hold");
```

`adjustThermostat(70, 25)` має повертати `"Heat: 7.0 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(70, 25), "Heat: 7.0 degrees Fahrenheit");
```

`adjustThermostat(72, 18)` має повертати `"Cool: 7.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(72, 18), "Cool: 7.6 degrees Fahrenheit");
```

`adjustThermostat(212, 100)` має повертати `"Hold"`.

```js
assert.equal(adjustThermostat(212, 100), "Hold");
```

`adjustThermostat(59, 22)` має повертати `"Heat: 12.6 degrees Fahrenheit"`.

```js
assert.equal(adjustThermostat(59, 22), "Heat: 12.6 degrees Fahrenheit");
```

# --seed--

## --seed-contents--

```js
function adjustThermostat(currentF, targetC) {

  return currentF;
}
```

# --solutions--

```js
function adjustThermostat(currentF, targetC) {
  const targetF = targetC * 1.8 + 32;
  const diff = Math.abs(targetF - currentF).toFixed(1);

  if (currentF < targetF) {
    return `Heat: ${diff} degrees Fahrenheit`;
  } else if (currentF > targetF) {
    return `Cool: ${diff} degrees Fahrenheit`;
  } else {
    return "Hold"
  }
}
```
