---
id: 68cae5b538ff798bbd4da006
title: "Завдання 68: маскувальник кредитної картки"
challengeType: 28
dashedName: challenge-68
---

# --description--

Дано рядок із номерами кредитних карток, поверніть його замасковану версію з урахуванням таких обмежень:

- Рядок міститиме чотири множини по чотири цифри (`0-9`), причому всі множини розділені одним пробілом або одним дефісом (`-`).
- Замініть усі числа, крім останніх чотирьох, на зірочку (`*`).
- Решту символів залиште без змін.

Наприклад, для `"4012-8888-8888-1881"` поверніть `"****-****-****-1881"`.

# --hints--

`mask("4012-8888-8888-1881")` має повернути `"****-****-****-1881"`.

```js
assert.equal(mask("4012-8888-8888-1881"), "****-****-****-1881");
```

`mask("5105 1051 0510 5100")` має повернути `"**** **** **** 5100"`.

```js
assert.equal(mask("5105 1051 0510 5100"), "**** **** **** 5100");
```

`mask("6011 1111 1111 1117")` має повернути `"**** **** **** 1117"`.

```js
assert.equal(mask("6011 1111 1111 1117"), "**** **** **** 1117");
```

`mask("2223-0000-4845-0010")` має повернути `"****-****-****-0010"`.

```js
assert.equal(mask("2223-0000-4845-0010"), "****-****-****-0010");
```

# --seed--

## --seed-contents--

```js
function mask(card) {

  return card;
}
```

# --solutions--

```js
function mask(card) {
  const split = card.split(card.includes('-') ? '-' : ' ');
  
  return card.includes('-') ?
    `****-****-****-${split[3]}` :
    `**** **** **** ${split[3]}`;
}
```
