---
id: 68d2ba1468508398389487d0
title: "Завдання 75: прихований скарб"
challengeType: 28
dashedName: challenge-75
---

# --description--

Дано двовимірний масив, який представляє карту дна океану з прихованим скарбом, та масив із координатами (`[row, column]`) для наступного занурення у пошуках скарбу. Поверніть `"Empty"`, `"Found"` або `"Recovered"` за такими правилами:

- Даний двовимірний масив міститиме рівно один не знайдений скарб, який займає кілька клітинок.
- Кожна клітинка в двовимірному масиві міститиме одне з наступних значень:
  - `"-"`: немає скарбу.
  - `"O"`: частина скарбу, яка ще не знайдена.
  - `"X"`: частина скарбу, яка вже знайдена.

- Якщо місце занурення не містить скарбу, поверніть `"Empty"`.
- Якщо місце занурення знаходить скарб, але принаймні одна інша частина скарбу ще не знайдена, поверніть `"Found"`.
- Якщо місце занурення знаходить останню не знайдену частину скарбу, поверніть `"Recovered"`.

Наприклад, якщо дано:

```json
[
  [ "-", "X"],
  [ "-", "X"],
  [ "-", "O"]
]
```

І `[2, 1]` як координати місця занурення, поверніть `"Recovered"`, бо занурення знайшло останню не знайдену частину скарбу.

# --hints--

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1])` має повернути `"Recovered"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 1]), "Recovered");
```

`dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0])` має повернути `"Empty"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "X"], [ "-", "O"]], [2, 0]), "Empty");
```

`dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1])` має повернути `"Found"`.

```js
assert.equal(dive([[ "-", "X"], [ "-", "O"], [ "-", "O"]], [1, 1]), "Found");
```

`dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2])` має повернути `"Found"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "X", "O", "X"], [ "-", "-", "-"]], [1, 2]), "Found");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0])` має повернути `"Recovered"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [2, 0]), "Recovered");
```

`dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2])` має повернути `"Empty"`.

```js
assert.equal(dive([[ "-", "-", "-"], [ "-", "-", "-"], [ "O", "X", "X"]], [1, 2]), "Empty");
```

# --seed--

## --seed-contents--

```js
function dive(map, coordinates) {

  return map;
}
```

# --solutions--

```js
function dive(map, coordinates) {
  const [row, col] = coordinates;
  const target = map[row][col];

  if (target === "-") {
    return "Empty";
  }

  if (target === "O") {
    for (let r = 0; r < map.length; r++) {
      for (let c = 0; c < map[r].length; c++) {
        if ((r !== row || c !== col) && map[r][c] === "O") {
          return "Found";
        }
      }
    }
    return "Recovered";
  }

  return "Found";
}
```
