---
id: 68d2ba1468508398389487ce
title: "Завдання 73: Говори мудро, ти повинен"
challengeType: 28
dashedName: challenge-73
---

# --description--

Дано речення, поверніть його версію, яка звучить як порада від мудрого вчителя, використовуючи такі правила:

- Слова розділені одним пробілом.
- Знайдіть перше входження одного з наступних слів у реченні: `"have"`, `"must"`, `"are"`, `"will"`, `"can"`.
- Перемістіть усі слова до і включно з цим словом у кінець речення і:
  - Збережіть порядок слів при переміщенні.
  - Зробіть їх усі у нижньому регістрі.
  - Додайте перед ними кому та пробіл.
- Зробіть першу літеру нового першого слова речення великою.
- Усі задані речення закінчуються одним розділовим знаком. Збережіть оригінальну пунктуацію речення і перемістіть її в кінець нового речення.
- Поверніть нове речення, переконайтеся, що між словами є лише один пробіл і немає пробілів на початку чи в кінці речення.

Наприклад, для `"You must speak wisely."` поверніть `"Speak wisely, you must."`

# --hints--

`wiseSpeak("You must speak wisely.")` має повернути `"Speak wisely, you must."`

```js
assert.equal(wiseSpeak("You must speak wisely."), "Speak wisely, you must.");
```

`wiseSpeak("You can do it!")` має повернути `"Do it, you can!"`

```js
assert.equal(wiseSpeak("You can do it!"), "Do it, you can!");
```

`wiseSpeak("Do you think you will complete this?")` має повернути `"Complete this, do you think you will?"`

```js
assert.equal(wiseSpeak("Do you think you will complete this?"), "Complete this, do you think you will?");
```

`wiseSpeak("All your base are belong to us.")` має повернути `"Belong to us, all your base are."`

```js
assert.equal(wiseSpeak("All your base are belong to us."), "Belong to us, all your base are.");
```

`wiseSpeak("You have much to learn.")` має повернути `"Much to learn, you have."`

```js
assert.equal(wiseSpeak("You have much to learn."), "Much to learn, you have.");
```

# --seed--

## --seed-contents--

```js
function wiseSpeak(sentence) {

  return sentence;
}
```

# --solutions--

```js
function wiseSpeak(sentence) {
  const triggers = ["have", "must", "are", "will", "can"];
  const punctuation = sentence[sentence.length - 1];
  const words = sentence.split(" ");

  const triggerIndex = words.findIndex(w => triggers.includes(w));

  const toMove = words.slice(0, triggerIndex + 1).map(w => w.toLowerCase());
  const remaining = words.slice(triggerIndex + 1);

  const newStart = remaining.map((w, i) => {
    return i === 0 ? w[0].toUpperCase() + w.slice(1) :
      i === remaining.length - 1 ? w.slice(0, w.length - 1) : w;
  })

  return newStart.join(" ") + ", " + toMove.join(" ") + punctuation;
}
```
