---
id: 68d30845cc08266018fc46be
title: "Завдання 78: послідовність цілих чисел"
challengeType: 28
dashedName: challenge-78
---

# --description--

Дано додатне ціле число. Поверніть рядок, що містить усі цілі числа від `1` до заданого числа включно, у числовому порядку.

Наприклад, для `5` поверніть `"12345"`.

# --hints--

`sequence(5)` має повернути `"12345"`.

```js
assert.equal(sequence(5), "12345");
```

`sequence(10)` має повернути `"12345678910"`.

```js
assert.equal(sequence(10), "12345678910");
```

`sequence(1)` має повернути `"1"`.

```js
assert.strictEqual(sequence(1), "1");
```

`sequence(27)` має повернути `"123456789101112131415161718192021222324252627"`.

```js
assert.equal(sequence(27), "123456789101112131415161718192021222324252627");
```

# --seed--

## --seed-contents--

```js
function sequence(n) {

  return n;
}
```

# --solutions--

```js
function sequence(n) {
  let result = "";
  for (let i = 1; i <= n; i++) {
    result += i;
  }
  return result;
}
```
