---
id: 68d30845cc08266018fc46bc
title: "Завдання 76: доповнююча ДНК"
challengeType: 28
dashedName: challenge-76
---

# --description--

Дано рядок, що представляє послідовність ДНК, поверніть її доповнюючу нитку за такими правилами:

- ДНК складається з літер `"A"`, `"C"`, `"G"` та `"T"`.
- Літери `"A"` та `"T"` доповнюють одна одну.
- Літери `"C"` та `"G"` доповнюють одна одну.

Наприклад, для `"ACGT"` поверніть `"TGCA"`.

# --hints--

`complementaryDNA("ACGT")` має повернути `"TGCA"`.

```js
assert.equal(complementaryDNA("ACGT"), "TGCA");
```

`complementaryDNA("ATGCGTACGTTAGC")` має повернути `"TACGCATGCAATCG"`.

```js
assert.equal(complementaryDNA("ATGCGTACGTTAGC"), "TACGCATGCAATCG");
```

`complementaryDNA("GGCTTACGATCGAAG")` має повернути `"CCGAATGCTAGCTTC"`.

```js
assert.equal(complementaryDNA("GGCTTACGATCGAAG"), "CCGAATGCTAGCTTC");
```

`complementaryDNA("GATCTAGCTAGGCTAGCTAG")` має повернути `"CTAGATCGATCCGATCGATC"`.

```js
assert.equal(complementaryDNA("GATCTAGCTAGGCTAGCTAG"), "CTAGATCGATCCGATCGATC");
```

# --seed--

## --seed-contents--

```js
function complementaryDNA(strand) {

  return strand;
}
```

# --solutions--

```js
function complementaryDNA(strand) {
  const complements = { A: "T", T: "A", C: "G", G: "C" };
  return strand
    .split("")
    .map(c => complements[c])
    .join("");
}
```
