---
id: 68d30fc57588d97fd3027b30
title: "Завдання 79: Navigator"
challengeType: 28
dashedName: challenge-79
---

# --description--

28 жовтня 1994 року вийшов Netscape Navigator, який допоміг мільйонам досліджувати ранній веб.

Дано масив команд браузера, які ви виконали у Netscape Navigator. Поверніть поточну сторінку, на якій ви перебуваєте після виконання всіх команд за такими правилами:

- Ви завжди починаєте зі сторінки `"Home"`, яка не входить до масиву команд.
- Дійсні команди:
  - `"Visit Page"`: де `"Page"` — це назва сторінки, яку ви відвідуєте. Наприклад, `"Visit About"` переводить вас на сторінку `"About"`. Коли ви переходите на нову сторінку, обов’язково скиньте будь-яку історію вперед.
  - `"Back"`: повертає вас на попередню сторінку в історії або залишає на поточній, якщо попередньої немає.
  - `"Forward"`: переходить вперед в історії на сторінку, з якої ви прийшли, або залишає на поточній, якщо такої немає.

Наприклад, для `["Visit About Us", "Back", "Forward"]` поверніть `"About Us"`.

# --hints--

`navigate(["Visit About Us", "Back", "Forward"])` має повернути `"About Us"`.

```js
assert.equal(navigate(["Visit About Us", "Back", "Forward"]), "About Us");
```

`navigate(["Forward"])` має повернути `"Home"`.

```js
assert.equal(navigate(["Forward"]), "Home");
```

`navigate(["Back"])` має повернути `"Home"`.

```js
assert.equal(navigate(["Back"]), "Home");
```

`navigate(["Visit About Us", "Visit Gallery"])` має повернути `"Gallery"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery"]), "Gallery");
```

`navigate(["Visit About Us", "Visit Gallery", "Back", "Back"])` має повернути `"Home"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Gallery", "Back", "Back"]), "Home");
```

`navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"])` має повернути `"Contact"`.

```js
assert.equal(navigate(["Visit About", "Visit Gallery", "Back", "Visit Contact", "Forward"]), "Contact");
```

`navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"])` має повернути `"Visit Us"`.

```js
assert.equal(navigate(["Visit About Us", "Visit Visit Us", "Forward", "Visit Contact Us", "Back"]), "Visit Us");
```

# --seed--

## --seed-contents--

```js
function navigate(commands) {

  return commands;
}
```

# --solutions--

```js
function navigate(commands) {
  const history = ["Home"];
  let currentPageIndex = 0;

  for (const command of commands) {
    if (command.startsWith("Visit")) {
      history.splice(currentPageIndex + 1);
      history.push(command.slice(6));
      currentPageIndex++;
    } else if(command === "Back" && currentPageIndex > 0) {
      currentPageIndex--;
    } else if (command === "Forward" && currentPageIndex < history.length - 1) {
      currentPageIndex++;
    }
  }

  return history[currentPageIndex];
}
```
