---
id: 68d30845cc08266018fc46bd
title: "Завдання 77: форматування тривалості"
challengeType: 28
dashedName: challenge-77
---

# --description--

Дано ціле число секунд, поверніть рядок, який представляє ту саму тривалість у форматі `"H:MM:SS"`, де `"H"` — це кількість годин, `"MM"` — кількість хвилин, а `"SS"` — кількість секунд. Поверніть час, використовуючи такі правила:

- Секунди: завжди мають бути двозначними.
- Хвилини: слід опускати провідні нулі, якщо вони не потрібні. Використовуйте `"0"`, якщо тривалість менша за одну хвилину.
- Години: слід включати лише якщо вони більші за нуль.

# --hints--

`format(500)` має повертати `"8:20"`.

```js
assert.equal(format(500), "8:20");
```

`format(4000)` має повертати `"1:06:40"`.

```js
assert.equal(format(4000), "1:06:40");
```

`format(1)` має повертати `"0:01"`.

```js
assert.equal(format(1), "0:01");
```

`format(5555)` має повертати `"1:32:35"`.

```js
assert.equal(format(5555), "1:32:35");
```

`format(99999)` має повертати `"27:46:39"`.

```js
assert.equal(format(99999), "27:46:39");
```

# --seed--

## --seed-contents--

```js
function format(seconds) {

  return seconds;
}
```

# --solutions--

```js
function format(seconds) {
  const h = Math.floor(seconds / 3600);
  const m = Math.floor((seconds % 3600) / 60);
  const s = seconds % 60;

  const secondsStr = s.toString().padStart(2, "0");
  const minutesStr = m.toString();

  if (h > 0) {
    const hoursStr = h.toString();
    return `${hoursStr}:${minutesStr.padStart(2, "0")}:${secondsStr}`;
  } else {
    return `${minutesStr}:${secondsStr}`;
  }
}
```
