---
id: 68e39ed6106dac2f0a98fd62
title: "Завдання 80: сортувальник е-пошти"
challengeType: 28
dashedName: challenge-80
---

# --description--

29 жовтня 1971 року було надіслано перший у світі лист електронної пошти, який представив формат `username@domain`, що ми досі використовуємо. Зараз існують мільярди адрес е-пошти.

У цьому завданні вам дають список адрес е-пошти, які потрібно відсортувати спочатку за алфавітом за доменним іменем (частина після `@`), а потім за іменем користувача (частина перед `@`).

- Сортування має бути нечутливим до регістру.
- Якщо кілька адрес мають однаковий домен, сортуйте їх за іменем користувача.
- Поверніть масив відсортованих адрес.
- Повернені адреси мають зберігати свій оригінальний регістр.

Наприклад, для `["jill@mail.com", "john@example.com", "jane@example.com"]` поверніть `["jane@example.com", "john@example.com", "jill@mail.com"]`.

# --hints--

`sort(["jill@mail.com", "john@example.com", "jane@example.com"])` має повернути `["jane@example.com", "john@example.com", "jill@mail.com"]`.

```js
assert.deepEqual(sort(["jill@mail.com", "john@example.com", "jane@example.com"]), ["jane@example.com", "john@example.com", "jill@mail.com"]);
```

`sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"])` має повернути `["bob@mail.com", "carol@mail.com", "alice@zoo.com"]`.

```js
assert.deepEqual(sort(["bob@mail.com", "alice@zoo.com", "carol@mail.com"]), ["bob@mail.com", "carol@mail.com", "alice@zoo.com"]);
```

`sort(["user@z.com", "user@y.com", "user@x.com"])` має повернути `["user@x.com", "user@y.com", "user@z.com"]`.

```js
assert.deepEqual(sort(["user@z.com", "user@y.com", "user@x.com"]), ["user@x.com", "user@y.com", "user@z.com"]);
```

`sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"])` має повернути `["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]`.

```js
assert.deepEqual(sort(["sam@MAIL.com", "amy@mail.COM", "bob@Mail.com"]), ["amy@mail.COM", "bob@Mail.com", "sam@MAIL.com"]);
```

`sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"])` має повернути `["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]`.

```js
assert.deepEqual(sort(["simon@beta.com", "sammy@alpha.com", "Sarah@Alpha.com", "SAM@ALPHA.com", "Simone@Beta.com", "sara@alpha.com"]), ["SAM@ALPHA.com", "sammy@alpha.com", "sara@alpha.com", "Sarah@Alpha.com", "simon@beta.com", "Simone@Beta.com"]);
```

# --seed--

## --seed-contents--

```js
function sort(emails) {

  return emails;
}
```

# --solutions--

```js
function sort(emails) {
  return emails.slice().sort((a, b) => {
    const [userA, domainA] = a.split('@');
    const [userB, domainB] = b.split('@');

    const domainCompare = domainA.toLowerCase().localeCompare(domainB.toLowerCase());
    if (domainCompare !== 0) return domainCompare;

    return userA.toLowerCase().localeCompare(userB.toLowerCase());
  });
}
```
