---
id: 68e39ed6106dac2f0a98fd63
title: "Завдання 81: n-те просте число"
challengeType: 28
dashedName: challenge-81
---

# --description--

Просте число — це додатне ціле число більше за 1, яке ділиться лише на 1 та саме на себе. Перші п’ять простих чисел — `2`, `3`, `5`, `7` та `11`.

Дано додатне ціле число `n`, поверніть `n`-те просте число. Наприклад, якщо дано `5`, поверніть 5-те просте число: `11`.

# --hints--

`nthPrime(5)` має повернути `11`.

```js
assert.equal(nthPrime(5), 11);
```

`nthPrime(10)` має повернути `29`.

```js
assert.equal(nthPrime(10), 29);
```

`nthPrime(16)` має повернути `53`.

```js
assert.equal(nthPrime(16), 53);
```

`nthPrime(99)` має повернути `523`.

```js
assert.equal(nthPrime(99), 523);
```

`nthPrime(1000)` має повернути `7919`.

```js
assert.equal(nthPrime(1000), 7919);
```

# --seed--

## --seed-contents--

```js
function nthPrime(n) {

  return n;
}
```

# --solutions--

```js
function nthPrime(n) {
  const primes = [];
  let num = 2;

  while (primes.length < n) {
    let isPrime = true;
    for (let i = 2; i * i <= num; i++) {
      if (num % i === 0) {
        isPrime = false;
        break;
      }
    }
    if (isPrime) primes.push(num);
    num++;
  }

  return primes[n - 1];
}
```
