---
id: 68e39ed6106dac2f0a98fd65
title: "Завдання 83: валідація підпису"
challengeType: 28
dashedName: challenge-83
---

# --description--

Дано рядок повідомлення, рядок секретного ключа та число підпису. Визначте, чи є підпис дійсним, використовуючи цей метод кодування:

- Літери в повідомленні та секретному ключі мають такі значення:
  - `a` до `z` мають відповідні значення `1` до `26`.
  - `A` до `Z` мають відповідні значення `27` до `52`.
- Всі інші символи не мають значення.
- Обчисліть підпис, взявши суму значень повідомлення та суму значень секретного ключа.

Наприклад, для повідомлення `"foo"` і секретного ключа `"bar"` підпис буде `57`:

```md
f (6) + o (15) + o (15) = 36
b (2) + a (1) + r (18) = 21
36 + 21 = 57
```

Перевірте, чи збігається обчислений підпис із наданим підписом.

# --hints--

`verify("foo", "bar", 57)` має повернути `true`.

```js
assert.isTrue(verify("foo", "bar", 57));
```

`verify("foo", "bar", 54)` має повернути `false`.

```js
assert.isFalse(verify("foo", "bar", 54));
```

`verify("freeCodeCamp", "Rocks", 238)` має повернути `true`.

```js
assert.isTrue(verify("freeCodeCamp", "Rocks", 238));
```

`verify("Is this valid?", "No", 210)` має повернути `false`.

```js
assert.isFalse(verify("Is this valid?", "No", 210));
```

`verify("Is this valid?", "Yes", 233)` має повернути `true`.

```js
assert.isTrue(verify("Is this valid?", "Yes", 233));
```

`verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514)` має повернути `true`.

```js
assert.isTrue(verify("Check out the freeCodeCamp podcast,", "in the mobile app", 514));
```

# --seed--

## --seed-contents--

```js
function verify(message, key, signature) {

  return message;
}
```

# --solutions--

```js
function verify(message, key, signature) {
  function charValue(ch) {
    if (ch >= 'a' && ch <= 'z') return ch.charCodeAt(0) - 'a'.charCodeAt(0) + 1;
    if (ch >= 'A' && ch <= 'Z') return ch.charCodeAt(0) - 'A'.charCodeAt(0) + 27;
    return 0;
  }

  function computeSum(str) {
    let sum = 0;
    for (let ch of str) {
      sum += charValue(ch);
    }
    return sum;
  }

  const total = computeSum(message) + computeSum(key);
  return total === signature;
}
```
