---
id: 68ee9e3066cfd4eb2328e8a5
title: "Завдання 86: пошук зображень"
challengeType: 28
dashedName: challenge-86
---

# --description--

4 листопада 2001 року Google запустив пошук зображень, що дозволяє людям знаходити зображення за пошуковими термінами. У цьому завданні ви імітуєте пошук зображень.

Дано масив назв зображень і пошуковий термін, поверніть масив назв зображень, які містять цей пошуковий термін.

- Ігноруйте регістр при порівнянні пошукових термінів.
- Поверніть зображення в тому ж порядку, в якому вони з’являються у вхідному масиві.

# --hints--

`imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog")` має повертати `["dog.png"]`.

```js
assert.deepEqual(imageSearch(["dog.png", "cat.jpg", "parrot.jpeg"], "dog"), ["dog.png"]);
```

`imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun")` має повертати `["Sunset.jpg", "sunflower.jpeg"]`.

```js
assert.deepEqual(imageSearch(["Sunset.jpg", "Beach.png", "sunflower.jpeg"], "sun"), ["Sunset.jpg", "sunflower.jpeg"]);
```

`imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG")` має повертати `["Moon.png", "stars.png"]`.

```js
assert.deepEqual(imageSearch(["Moon.png", "sun.jpeg", "stars.png"], "PNG"), ["Moon.png", "stars.png"]);
```

`imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat")` має повертати `["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]`.

```js
assert.deepEqual(imageSearch(["cat.jpg", "dogToy.jpeg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"], "Cat"), ["cat.jpg", "kitty-cat.png", "catNip.jpeg", "franken_cat.gif"]);
```

# --seed--

## --seed-contents--

```js
function imageSearch(images, term) {

  return images;
}
```

# --solutions--

```js
function imageSearch(images, term) {
  const lowerTerm = term.toLowerCase();
  return images.filter(img => img.toLowerCase().includes(lowerTerm));
}
```
