---
id: 68ee9e3066cfd4eb2328e8a8
title: "Завдання 89: підрахунок карт"
challengeType: 28
dashedName: challenge-89
---

# --description--

Стандартна колода карт містить 13 унікальних карт у кожній масті. За заданим цілим числом, що представляє кількість карт, які потрібно вибрати з колоди, поверніть кількість унікальних комбінацій карт, які ви можете вибрати.

- Порядок не має значення. Вибір карти A, а потім карти B — це те саме, що вибір карти B, а потім карти A.

Наприклад, для `52` поверніть `1`. Існує лише одна комбінація з 52 карт, яку можна вибрати з колоди у 52 карти. А для `2` поверніть `1326`. Існує 1326 комбінацій карт, які ви можете отримати, вибираючи 2 карти з колоди.

# --hints--

`combinations(52)` має повернути `1`.

```js
assert.equal(combinations(52), 1);
```

`combinations(1)` має повернути `52`.

```js
assert.equal(combinations(1), 52);
```

`combinations(2)` має повернути `1326`.

```js
assert.equal(combinations(2), 1326);
```

`combinations(5)` має повернути `2598960`.

```js
assert.equal(combinations(5), 2598960);
```

`combinations(10)` має повернути `15820024220`.

```js
assert.equal(combinations(10), 15820024220);
```

`combinations(50)` має повернути `1326`.

```js
assert.equal(combinations(50), 1326);
```

# --seed--

## --seed-contents--

```js
function combinations(cards) {

  return cards;
}
```

# --solutions--

```js
function combinations(cards) {
  const n = 52;

  function factorial(x) {
    let result = 1;
    for (let i = 2; i <= x; i++) {
      result *= i;
    }
    return result;
  }

  return factorial(n) / (factorial(cards) * factorial(n - cards));
}
```
