---
id: 68ee9e3066cfd4eb2328e8a4
title: "Завдання 85: фішка для підрахунку слів"
challengeType: 28
dashedName: challenge-85
---

# --description--

Дано рядок речення, поверніть число слів, які є в реченні.

- Слова — це будь-яка послідовність символів, окрім пробілів, і вони розділені одним пробілом.

# --hints--

`countWords("Hello world")` має повернути `2`.

```js
assert.equal(countWords("Hello world"), 2);
```

`countWords("The quick brown fox jumps over the lazy dog.")` має повернути `9`.

```js
assert.equal(countWords("The quick brown fox jumps over the lazy dog."), 9);
```

`countWords("I like coding challenges!")` має повернути `4`.

```js
assert.equal(countWords("I like coding challenges!"), 4);
```

`countWords("Complete the challenge in JavaScript and Python.")` має повернути `7`.

```js
assert.equal(countWords("Complete the challenge in JavaScript and Python."), 7);
```

`countWords("The missing semi-colon crashed the entire internet.")` має повернути `7`.

```js
assert.equal(countWords("The missing semi-colon crashed the entire internet."), 7);
```

# --seed--

## --seed-contents--

```js
function countWords(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countWords(sentence) {

  return sentence.split(' ').length;
}
```
