---
id: 68f6587287ad1f4ad39b0c84
title: "Завдання 98: підрахунок прямокутників"
challengeType: 28
dashedName: challenge-98
---

# --description--

Дано два натуральні числа, які представляють ширину та висоту прямокутника. Визначте, скільки прямокутників можна розмістити всередині заданого.

- Рахуйте лише прямокутники з цілою шириною та висотою.

Наприклад, для `1` і `3` поверніть `6`. Три прямокутники 1x1, два прямокутники 1x2 і один прямокутник 1x3.

# --hints--

`countRectangles(1, 3)` має повернути `6`.

```js
assert.equal(countRectangles(1, 3), 6);
```

`countRectangles(3, 2)` має повернути `18`.

```js
assert.equal(countRectangles(3, 2), 18);
```

`countRectangles(1, 2)` має повернути `3`.

```js
assert.equal(countRectangles(1, 2), 3);
```

`countRectangles(5, 4)` має повернути `150`.

```js
assert.equal(countRectangles(5, 4), 150);
```

`countRectangles(11, 19)` має повернути `12540`.

```js
assert.equal(countRectangles(11, 19), 12540);
```

# --seed--

## --seed-contents--

```js
function countRectangles(width, height) {

  return width;
}
```

# --solutions--

```js
function countRectangles(width, height) {
  let count = 0;

  for (let w = 1; w <= width; w++) {
    for (let h = 1; h <= height; h++) {
      count += (width - w + 1) * (height - h + 1);
    }
  }

  return count;
}
```
