---
id: 68f6587287ad1f4ad39b0c81
title: "Завдання 95: Зсув масиву"
challengeType: 28
dashedName: challenge-95
---

# --description--

Дано масив і ціле число, яке вказує, на скільки позицій потрібно зсунути масив. Поверніть зсунений масив.

- Додатне ціле число зсуває масив вліво.
- Від’ємне ціле число зсуває масив вправо.
- Зсув відбувається циклічно, тобто елементи, що виходять за межі, повертаються з іншого боку масиву.

Наприклад, якщо дано `[1, 2, 3]` і `1`, зсуньте масив на 1 позицію вліво, повернувши `[2, 3, 1]`.

# --hints--

`shiftArray([1, 2, 3], 1)` має повернути `[2, 3, 1]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], 1), [2, 3, 1]);
```

`shiftArray([1, 2, 3], -1)` має повернути `[3, 1, 2]`.

```js
assert.deepEqual(shiftArray([1, 2, 3], -1), [3, 1, 2]);
```

`shiftArray(["alpha", "bravo", "charlie"], 5)` має повернути `["charlie", "alpha", "bravo"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], 5), ["charlie", "alpha", "bravo"]);
```

`shiftArray(["alpha", "bravo", "charlie"], -11)` має повернути `["bravo", "charlie", "alpha"]`.

```js
assert.deepEqual(shiftArray(["alpha", "bravo", "charlie"], -11), ["bravo", "charlie", "alpha"]);
```

`shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15)` має повернути `[5, 6, 7, 8, 9, 0, 1, 2, 3, 4]`.

```js
assert.deepEqual(shiftArray([0, 1, 2, 3, 4, 5, 6, 7, 8, 9], 15), [5, 6, 7, 8, 9, 0, 1, 2, 3, 4]);
```

# --seed--

## --seed-contents--

```js
function shiftArray(arr, n) {

  return arr;
}
```

# --solutions--

```js
function shiftArray(arr, n) {
  const len = arr.length;
  n = n % len;
  if (n < 0) n += len;

  return arr.slice(n).concat(arr.slice(0, n));
}
```
