---
id: 68f6587287ad1f4ad39b0c7d
title: "Завдання 91: Пошук слова"
challengeType: 28
dashedName: challenge-91
---

# --description--

Дано матрицю (масив масивів) з окремих літер і слово для пошуку, поверніть початкові та кінцеві індекси слова в матриці.

- Дана матриця буде заповнена всіма літерами в нижньому регістрі (`a-z`).
- Слово для пошуку завжди буде в матриці рівно один раз.
- Слово для пошуку завжди буде розташоване по прямій у одному з таких напрямків:
  - зліва направо
  - справа наліво
  - зверху вниз
  - знизу вгору

Наприклад, для матриці:

```md
[
  ["a", "c", "t"],
  ["t", "a", "t"],
  ["c", "t", "c"]
]
```

І слова `"cat"`, поверніть:

```md
[[0, 1], [2, 1]]
```

Де `[0, 1]` — це індекси для `"c"` (початок слова), а `[2, 1]` — індекси для `"t"` (кінець слова).

# --hints--

`findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat")` має повернути `[[0, 1], [2, 1]]`.

```js
assert.deepEqual(findWord([["a", "c", "t"], ["t", "a", "t"], ["c", "t", "c"]], "cat"), [[0, 1], [2, 1]]);
```

`findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog")` має повернути `[[0, 0], [0, 2]]`.

```js
assert.deepEqual(findWord([["d", "o", "g"], ["o", "g", "d"], ["d", "g", "o"]], "dog"), [[0, 0], [0, 2]]);
```

`findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish")` має повернути `[[3, 3], [0, 3]]`.

```js
assert.deepEqual(findWord([["h", "i", "s", "h"], ["i", "s", "f", "s"], ["f", "s", "i", "i"], ["s", "h", "i", "f"]], "fish"), [[3, 3], [0, 3]]);
```

`findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox")` має повернути `[[1, 3], [1, 1]]`.

```js
assert.deepEqual(findWord([["f", "x", "o", "x"], ["o", "x", "o", "f"], ["f", "o", "f", "x"], ["f", "x", "x", "o"]], "fox"), [[1, 3], [1, 1]]);
```

# --seed--

## --seed-contents--

```js
function findWord(matrix, word) {

  return matrix;
}
```

# --solutions--

```js
function findWord(matrix, word) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const len = word.length;

  const directions = [
    [0, 1],
    [0, -1],
    [1, 0],
    [-1, 0]
  ];

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      for (let [dr, dc] of directions) {
        let match = true;
        for (let i = 0; i < len; i++) {
          const nr = r + dr * i;
          const nc = c + dc * i;
          if (nr < 0 || nr >= rows || nc < 0 || nc >= cols || matrix[nr][nc] !== word[i]) {
            match = false;
            break;
          }
        }
        if (match) {
          return [
            [r, c],
            [r + dr * (len - 1), c + dc * (len - 1)]
          ];
        }
      }
    }
  }
}
```
