---
id: 68f6587287ad1f4ad39b0c82
title: "Завдання 96: Чи це вихідні?"
challengeType: 28
dashedName: challenge-96
---

# --description--

Дано дату у форматі `"YYYY-MM-DD"`, поверніть число днів, що залишилися до вихідних.

- Вихідні починаються в суботу.
- Якщо дана дата — субота або неділя, поверніть `"It's the weekend!"`.
- Інакше поверніть `"X days until the weekend."`, де `X` — це число днів до суботи.
- Якщо `X` дорівнює `1`, використовуйте `"day"` (однина) замість `"days"` (множина).
- Переконайтеся, що обчислення ігнорує ваш локальний часовий пояс.

# --hints--

`daysUntilWeekend("2025-11-14")` має повертати `"1 day until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-11-14"), "1 day until the weekend.");
```

`daysUntilWeekend("2025-01-01")` має повертати `"3 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2025-01-01"), "3 days until the weekend.");
```

`daysUntilWeekend("2025-12-06")` має повертати `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2025-12-06"), "It's the weekend!");
```

`daysUntilWeekend("2026-01-27")` має повертати `"4 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-01-27"), "4 days until the weekend.");
```

`daysUntilWeekend("2026-09-07")` має повертати `"5 days until the weekend."`.

```js
assert.equal(daysUntilWeekend("2026-09-07"), "5 days until the weekend.");
```

`daysUntilWeekend("2026-11-29")` має повертати `"It's the weekend!"`.

```js
assert.equal(daysUntilWeekend("2026-11-29"), "It's the weekend!");
```

# --seed--

## --seed-contents--

```js
function daysUntilWeekend(dateString) {

  return dateString;
}
```

# --solutions--

```js
function daysUntilWeekend(dateString) {
  const [year, month, day] = dateString.split("-").map(Number);
  const date = new Date(Date.UTC(year, month - 1, day));

  const dayOfWeek = date.getUTCDay();

  if (dayOfWeek === 6 || dayOfWeek === 0) {
    return "It's the weekend!";
  }

  const daysUntilSaturday = (6 - dayOfWeek);
  return `${daysUntilSaturday} day${daysUntilSaturday === 1 ? "" : "s"} until the weekend.`;
}
```
