---
id: 68f6587287ad1f4ad39b0c7e
title: "Завдання 92: витягнення розширення"
challengeType: 28
dashedName: challenge-92
---

# --description--

Дано рядок, що представляє ім’я файлу, поверніть розширення файлу.

- Розширення — це частина імені файлу, що йде після останньої крапки (`.`).
- Якщо ім’я файлу не містить крапки або закінчується крапкою, поверніть `"none"`.
- Розширення має бути повернене без змін, збереженням регістру.

# --hints--

`getExtension("document.txt")` має повернути `"txt"`.

```js
assert.equal(getExtension("document.txt"), "txt");
```

`getExtension("README")` має повернути `"none"`.

```js
assert.equal(getExtension("README"), "none");
```

`getExtension("image.PNG")` має повернути `"PNG"`.

```js
assert.equal(getExtension("image.PNG"), "PNG");
```

`getExtension(".gitignore")` має повернути `"gitignore"`.

```js
assert.equal(getExtension(".gitignore"), "gitignore");
```

`getExtension("archive.tar.gz")` має повернути `"gz"`.

```js
assert.equal(getExtension("archive.tar.gz"), "gz");
```

`getExtension("final.draft.")` має повернути `"none"`.

```js
assert.equal(getExtension("final.draft."), "none");
```

# --seed--

## --seed-contents--

```js
function getExtension(filename) {

  return filename;
}
```

# --solutions--

```js
function getExtension(filename) {
  const lastDot = filename.lastIndexOf('.');
  
  if (lastDot === -1 || lastDot === filename.length - 1) {
    return "none";
  }
  
  return filename.slice(lastDot + 1);
}
```
