---
id: 68f6587287ad1f4ad39b0c7f
title: "Завдання 93: Голосні та приголосні"
challengeType: 28
dashedName: challenge-93
---

# --description--

Дано рядок, поверніть масив із числом голосних і числом приголосних у цьому рядку.

- Голосні складаються з `a`, `e`, `i`, `o`, `u` у будь-якому регістрі.
- Приголосні — це всі інші літери у будь-якому регістрі.
- Ігноруйте всі символи, що не є літерами.

Наприклад, для `"Hello World"` поверніть `[3, 7]`.

# --hints--

`count("Hello World")` має повернути `[3, 7]`.

```js
assert.deepEqual(count("Hello World"), [3, 7]);
```

`count("JavaScript")` має повернути `[3, 7]`.

```js
assert.deepEqual(count("JavaScript"), [3, 7]);
```

`count("Python")` має повернути `[1, 5]`.

```js
assert.deepEqual(count("Python"), [1, 5]);
```

`count("freeCodeCamp")` має повернути `[5, 7]`.

```js
assert.deepEqual(count("freeCodeCamp"), [5, 7]);
```

`count("Hello, World!")` має повернути `[3, 7]`.

```js
assert.deepEqual(count("Hello, World!"), [3, 7]);
```

`count("The quick brown fox jumps over the lazy dog.")` має повернути `[11, 24]`.

```js
assert.deepEqual(count("The quick brown fox jumps over the lazy dog."), [11, 24]);
```

# --seed--

## --seed-contents--

```js
function count(str) {

  return str;
}
```

# --solutions--

```js
function count(str) {
  const vowels = 'aeiou';
  const consonants = 'bcdfghjklmnpqrstvwxyz';
  let v = 0, c = 0;

  for (let i=0; i<str.length; i++) {
    if (vowels.includes(str[i].toLowerCase())) {
      v++;
    }
    if (consonants.includes(str[i].toLowerCase())) {
      c++;
    }
  }

  return [v, c];
}
```
