---
id: 68f6587287ad1f4ad39b0c80
title: "Завдання 94: генератор підпису для е-пошти"
challengeType: 28
dashedName: challenge-94
---

# --description--

За заданими рядками з іменем, посадою та компанією людини поверніть підпис для е-пошти у вигляді одного рядка, використовуючи такі правила:

- Ім’я має йти першим, перед яким слід додати префікс, що залежить від першої літери імені. Для імен, що починаються (незалежно від регістру) з:
  - `A-I`: використовуйте `>>` як префікс.
  - `J-R`: використовуйте `--` як префікс.
  - `S-Z`: використовуйте `::` як префікс.
- Після імені має йти кома та пробіл (`, `).
- Посада та компанія мають йти після коми та пробілу, розділені `" at "` (з пробілами з обох боків).

Наприклад, за заданими `"Quinn Waverly"`, `"Founder and CEO"` та `"TechCo"` поверніть `"--Quinn Waverly, Founder and CEO at TechCo"`.

# --hints--

`generateSignature("Quinn Waverly", "Founder and CEO", "TechCo")` має повернути `"--Quinn Waverly, Founder and CEO at TechCo"`.

```js
assert.equal(generateSignature("Quinn Waverly", "Founder and CEO", "TechCo"), "--Quinn Waverly, Founder and CEO at TechCo");
```

`generateSignature("Alice Reed", "Engineer", "TechCo")` має повернути `">>Alice Reed, Engineer at TechCo"`.

```js
assert.equal(generateSignature("Alice Reed", "Engineer", "TechCo"), ">>Alice Reed, Engineer at TechCo");
```

`generateSignature("Tina Vaughn", "Developer", "example.com")` має повернути `"::Tina Vaughn, Developer at example.com"`.

```js
assert.equal(generateSignature("Tina Vaughn", "Developer", "example.com"), "::Tina Vaughn, Developer at example.com");
```

`generateSignature("B. B.", "Product Tester", "AcmeCorp")` має повернути `">>B. B., Product Tester at AcmeCorp"`.

```js
assert.equal(generateSignature("B. B.", "Product Tester", "AcmeCorp"), ">>B. B., Product Tester at AcmeCorp");
```

`generateSignature("windstorm", "Cloud Architect", "Atmospheronics")` має повернути `"::windstorm, Cloud Architect at Atmospheronics"`.

```js
assert.equal(generateSignature("windstorm", "Cloud Architect", "Atmospheronics"), "::windstorm, Cloud Architect at Atmospheronics");
```

# --seed--

## --seed-contents--

```js
function generateSignature(name, title, company) {

  return name;
}
```

# --solutions--

```js
function generateSignature(name, title, company) {
  const firstLetter = name[0].toUpperCase();
  let prefix;
  if ("ABCDEFGHI".includes(firstLetter)) {
    prefix = ">>";
  } else if ("JKLMNOPQR".includes(firstLetter)) {
    prefix = "--";
  } else {
    prefix = "::";
  }

  return `${prefix}${name}, ${title} at ${company}`;
}
```
