---
id: 68f6587287ad1f4ad39b0c83
title: "Завдання 97: найбільший спільний дільник"
challengeType: 28
dashedName: challenge-97
---

# --description--

Дано два натуральні числа, поверніть їхній найбільший спільний дільник (НСД).

- НСД двох цілих чисел — це найбільше число, яке ділить обидва числа без остачі.

Наприклад, дільниками `4` є `1`, `2` та `4`. Дільниками `6` є `1`, `2`, `3` та `6`. Тож, маючи `4` і `6`, поверніть `2`, найбільше число, яке зустрічається в обох множинах дільників.

# --hints--

`gcd(4, 6)` має повернути `2`.

```js
assert.equal(gcd(4, 6), 2);
```

`gcd(20, 15)` має повернути `5`.

```js
assert.equal(gcd(20, 15), 5);
```

`gcd(13, 17)` має повернути `1`.

```js
assert.equal(gcd(13, 17), 1);
```

`gcd(654, 456)` має повернути `6`.

```js
assert.equal(gcd(654, 456), 6);
```

`gcd(3456, 4320)` має повернути `864`.

```js
assert.equal(gcd(3456, 4320), 864);
```

# --seed--

## --seed-contents--

```js
function gcd(x, y) {

  return x;
}
```

# --solutions--

```js
function gcd(x, y) {
  while (y !== 0) {
    const temp = y;
    y = x % y;
    x = temp;
  }
  return x;
}
```
