---
id: 68ffb91507a5b645769328c4
title: "Завдання 101: конвертер заголовків markdown"
challengeType: 28
dashedName: challenge-101
---

# --description--

Дано рядок, який представляє заголовок у Markdown. Поверніть еквівалентний HTML-заголовок.

Дійсний заголовок Markdown повинен:

- Починатися з нуля або більше пробілів, за якими йдуть
- Від 1 до 6 символів решітки (`#`) підряд, потім
- Принаймні один пробіл. І нарешті,
- Текст заголовка.

Кількість символів решітки визначає рівень заголовка. Наприклад, один символ решітки відповідає тегу `h1`, а шість символів решітки — тегу `h6`.

Якщо заданий рядок не має точно такого формату, поверніть `"Invalid format"`.

Наприклад, для `"# My level 1 heading"` поверніть `"<h1>My level 1 heading</h1>"`.

Примітка: консоль може не відображати HTML-теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`convert("# My level 1 heading")` має повернути `"<h1>My level 1 heading</h1>"`.

```js
assert.equal(convert("# My level 1 heading"), "<h1>My level 1 heading</h1>");
```

`convert("My heading")` має повернути `"Invalid format"`.

```js
assert.equal(convert("My heading"), "Invalid format");
```

`convert("##### My level 5 heading")` має повернути `"<h5>My level 5 heading</h5>"`.

```js
assert.equal(convert("##### My level 5 heading"), "<h5>My level 5 heading</h5>");
```

`convert("#My heading")` має повернути `"Invalid format"`.

```js
assert.equal(convert("#My heading"), "Invalid format");
```

`convert("  ###  My level 3 heading")` має повернути `"<h3>My level 3 heading</h3>"`.

```js
assert.equal(convert("  ###  My level 3 heading"), "<h3>My level 3 heading</h3>");
```

`convert("####### My level 7 heading")` має повернути `"Invalid format"`.

```js
assert.equal(convert("####### My level 7 heading"), "Invalid format");
```

`convert("## My #2 heading")` має повернути `"<h2>My #2 heading</h2>"`.

```js
assert.equal(convert("## My #2 heading"), "<h2>My #2 heading</h2>");
```

# --seed--

## --seed-contents--

```js
function convert(heading) {

  return heading;
}
```

# --solutions--

```js
function convert(heading) {
  const trimmed = heading.trim()

  if (!/^#{1,6}\s+\S+/.test(trimmed)) return "Invalid format"

  const hashCount = trimmed.indexOf(' ');
  const text = trimmed.slice(hashCount).trimStart();

  return `<h${hashCount}>${text}</h${hashCount}>`;
}
```
