---
id: 68ffb91507a5b645769328cb
title: "Завдання 108: BuzzFizz"
challengeType: 28
dashedName: challenge-108
---

# --description--

Дано масив, визначте, чи є він правильною послідовністю FizzBuzz від 1 до останнього елемента масиву. Послідовність правильна, якщо:

- Числа, кратні 3, замінені на "Fizz"
- Числа, кратні 5, замінені на "Buzz"
- Числа, кратні і 3, і 5, замінені на "FizzBuzz"
- Всі інші числа залишаються цілими числами у зростаючому порядку, починаючи з 1.
- Масив має починатися з 1 і не містити пропущених або зайвих елементів.

# --hints--

`isFizzBuzz([1, 2, "Fizz", 4])` має повертати `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4]));
```

`isFizzBuzz([1, 2, 3, 4])` має повертати `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, 3, 4]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7])` має повертати `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", 7]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"])` має повертати `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "FizzBuzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"])` має повертати `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Fizz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"])` має повертати `false`.

```js
assert.isFalse(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, "Buzz"]));
```

`isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"])` має повертати `true`.

```js
assert.isTrue(isFizzBuzz([1, 2, "Fizz", 4, "Buzz", "Fizz", 7, 8, "Fizz", "Buzz", 11, "Fizz", 13, 14, "FizzBuzz", 16, 17, "Fizz", 19, "Buzz", "Fizz", 22, 23, "Fizz", "Buzz", 26, "Fizz", 28, 29, "FizzBuzz", 31, 32, "Fizz", 34, "Buzz", "Fizz", 37, 38, "Fizz", "Buzz", 41, "Fizz", 43, 44, "FizzBuzz", 46, 47, "Fizz", 49, "Buzz"]));
```

# --seed--

## --seed-contents--

```js
function isFizzBuzz(sequence) {

  return sequence;
}
```

# --solutions--

```js
function isFizzBuzz(sequence) {
  for (let i = 0; i < sequence.length; i++) {
    const n = i + 1;
    if (n % 15 === 0 && sequence[i] !== "FizzBuzz") return false;
    if (n % 3 === 0 && n % 5 !== 0 && sequence[i] !== "Fizz") return false;
    if (n % 5 === 0 && n % 3 !== 0 && sequence[i] !== "Buzz") return false;
    if (n % 3 !== 0 && n % 5 !== 0 && sequence[i] !== n) return false;
  }
  return true;
}
```
