---
id: 68ffb91507a5b645769328c5
title: "Завдання 102: найдовше слово"
challengeType: 28
dashedName: challenge-102
---

# --description--

Дано рядок із реченням, поверніть найдовше слово в цьому реченні.

- Слова розділені одним пробілом.
- До довжини слова враховуються лише літери (`a-z`, без урахування регістру).
- Якщо є кілька слів однакової довжини, поверніть перше, яке зустрічається.
- Поверніть слово так, як воно з’являється у заданому рядку, без розділових знаків.

# --hints--

`longestWord("The quick red fox")` має повернути `"quick"`.

```js
assert.equal(longestWord("The quick red fox"), "quick");
```

`longestWord("Hello coding challenge.")` має повернути `"challenge"`.

```js
assert.equal(longestWord("Hello coding challenge."), "challenge");
```

`longestWord("Do Try This At Home.")` має повернути `"This"`.

```js
assert.equal(longestWord("Do Try This At Home."), "This");
```

`longestWord("This sentence... has commas, ellipses, and an exclamation point!")` має повернути `"exclamation"`.

```js
assert.equal(longestWord("This sentence... has commas, ellipses, and an exclamation point!"), "exclamation");
```

`longestWord("A tie? No way!")` має повернути `"tie"`.

```js
assert.equal(longestWord("A tie? No way!"), "tie");
```

`longestWord("Wouldn't you like to know.")` має повернути `"Wouldnt"`.

```js
assert.equal(longestWord("Wouldn't you like to know."), "Wouldnt");
```

# --seed--

## --seed-contents--

```js
function longestWord(sentence) {

  return sentence;
}
```

# --solutions--

```js
function longestWord(sentence) {
  const words = sentence.split(" ");
  let longest = "";

  for (let word of words) {
    const cleaned = word.replace(/[^a-z]/gi, "");
    if (cleaned.length > longest.length) {
      longest = cleaned;
    }
  }

  return longest;
}
```
