---
id: 68ffb91507a5b645769328c9
title: "Завдання 106: валідатор повідомлень"
challengeType: 28
dashedName: challenge-106
---

# --description--

Дано рядок повідомлення та рядок валідації, визначте, чи є повідомлення дійсним.

- Повідомлення вважається дійсним, якщо кожне слово в повідомленні починається з відповідної літери в рядку валідації, у правильному порядку.
- Літери не чутливі до регістру.
- Слова в повідомленні розділені одним пробілом.

# --hints--

`isValidMessage("hello world", "hw")` має повертати `true`.

```js
assert.isTrue(isValidMessage("hello world", "hw"));
```

`isValidMessage("ALL CAPITAL LETTERS", "acl")` має повертати `true`.

```js
assert.isTrue(isValidMessage("ALL CAPITAL LETTERS", "acl"));
```

`isValidMessage("Coding challenge are boring.", "cca")` має повертати `false`.

```js
assert.isFalse(isValidMessage("Coding challenge are boring.", "cca"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD")` має повертати `true`.

```js
assert.isTrue(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLD"));
```

`isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT")` має повертати `false`.

```js
assert.isFalse(isValidMessage("The quick brown fox jumps over the lazy dog.", "TQBFJOTLDT"));
```

# --seed--

## --seed-contents--

```js
function isValidMessage(message, validator) {

  return message;
}
```

# --solutions--

```js
function isValidMessage(message, validation) {
  const words = message.split(" ");
  
  if (words.length !== validation.length) return false;

  for (let i = 0; i < words.length; i++) {
    if (words[i][0].toLowerCase() !== validation[i].toLowerCase()) {
      return false;
    }
  }

  return true;
}
```
