---
id: 68ffb91507a5b645769328c8
title: "Завдання 105: підрахунок символів"
challengeType: 28
dashedName: challenge-105
---

# --description--

Дано рядок речення, поверніть масив із підрахунком кожного символу в алфавітному порядку.

- Регістр літер не враховуйте — великі та малі літери вважайте однаковими.
- Ігноруйте числа, пробіли, розділові знаки тощо.
- Поверніть підрахунок і літеру у форматі `"letter count"`. Наприклад, `"a 3"`.
- Усі повернені літери мають бути в нижньому регістрі.
- Не повертайте підрахунок літер, яких немає у заданому рядку.

# --hints--

`countCharacters("hello world")` має повертати `["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]`.

```js
assert.deepEqual(countCharacters("hello world"), ["d 1", "e 1", "h 1", "l 3", "o 2", "r 1", "w 1"]);
```

`countCharacters("I love coding challenges!")` має повертати `["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]`.

```js
assert.deepEqual(countCharacters("I love coding challenges!"), ["a 1", "c 2", "d 1", "e 3", "g 2", "h 1", "i 2", "l 3", "n 2", "o 2", "s 1", "v 1"]);
```

`countCharacters("// TODO: Complete this challenge ASAP!")` має повертати `["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]`.

```js
assert.deepEqual(countCharacters("// TODO: Complete this challenge ASAP!"), ["a 3", "c 2", "d 1", "e 4", "g 1", "h 2", "i 1", "l 3", "m 1", "n 1", "o 3", "p 2", "s 2", "t 3"]);
```

# --seed--

## --seed-contents--

```js
function countCharacters(sentence) {

  return sentence;
}
```

# --solutions--

```js
function countCharacters(sentence) {
  const counts = {};

  for (let char of sentence.toLowerCase()) {
    if (char >= "a" && char <= "z") {
      counts[char] = (counts[char] || 0) + 1;
    }
  }

  return Object.keys(counts)
               .sort()
               .map(letter => `${letter} ${counts[letter]}`);
}
```
