---
id: 68ffb91507a5b645769328cc
title: "Завдання 109: Скільки мені років?"
challengeType: 28
dashedName: challenge-109
---

# --description--

За датою народження у форматі `YYYY-MM-DD` поверніть вік людини станом на 27 листопада 2025 року.

- Припускайте, що всі дати народження є дійсними і передують 27 листопада 2025 року.
- Поверніть вік у вигляді цілого числа.
- Обов’язково врахуйте, чи вже була у людини дата народження у 2025 році.

# --hints--

`calculateAge("2000-11-20")` має повертати `25`.

```js
assert.equal(calculateAge("2000-11-20"), 25);
```

`calculateAge("2000-12-01")` має повертати `24`.

```js
assert.equal(calculateAge("2000-12-01"), 24);
```

`calculateAge("2014-10-25")` має повертати `11`.

```js
assert.equal(calculateAge("2014-10-25"), 11);
```

`calculateAge("1994-01-06")` має повертати `31`.

```js
assert.equal(calculateAge("1994-01-06"), 31);
```

`calculateAge("1994-12-14")` має повертати `30`.

```js
assert.equal(calculateAge("1994-12-14"), 30);
```

# --seed--

## --seed-contents--

```js
function calculateAge(birthday) {

  return birthday;
}
```

# --solutions--

```js
function calculateAge(birthday) {
  const today = new Date("2025-11-27");
  const [year, month, day] = birthday.split("-").map(Number);
  const birthDate = new Date(year, month - 1, day);

  let age = today.getFullYear() - birthDate.getFullYear();

  if (
    today.getMonth() < birthDate.getMonth() ||
    (today.getMonth() === birthDate.getMonth() && today.getDate() < birthDate.getDate())
  ) {
    age--;
  }

  return age;
}
```
