---
id: 69162d64f96574d9bb629efc
title: "Завдання 111: траєкторія м’яча"
challengeType: 28
dashedName: challenge-111
---

# --description--

Сьогоднішнє завдання натхненне відеогрою Pong, яка вийшла 29 листопада 1972 року.

Дано матрицю (масив масивів), що містить розташування м’яча (`2`) і попереднє розташування м’яча (`1`). Потрібно повернути індекси матриці для наступного розташування м’яча.

- М’яч завжди рухається по прямій лінії.
- Напрямок руху визначається тим, як м’яч рухався від `1` до `2`.
- Краї матриці вважаються стінами. Якщо м’яч вдаряється об:
  - верхню або нижню стіну, він відскакує, змінюючи вертикальний напрямок на протилежний.
  - ліву або праву стіну, він відскакує, змінюючи горизонтальний напрямок на протилежний.
  - кут, він відскакує, змінюючи обидва напрямки.

# --hints--

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]])` має повернути `[2, 3]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,1,2,0], [0,0,0,0]]), [2, 3]);
```

`getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]])` має повернути `[3, 0]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,1,0], [0,2,0,0], [0,0,0,0]]), [3, 0]);
```

`getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]])` має повернути `[1, 2]`.

```js
assert.deepEqual(getNextLocation([[0,2,0,0], [1,0,0,0], [0,0,0,0], [0,0,0,0]]), [1, 2]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]])` має повернути `[1, 1]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [2,0,0,0], [0,1,0,0]]), [1, 1]);
```

`getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]])` має повернути `[2, 2]`.

```js
assert.deepEqual(getNextLocation([[0,0,0,0], [0,0,0,0], [0,0,1,0], [0,0,0,2]]), [2, 2]);
```

# --seed--

## --seed-contents--

```js
function getNextLocation(matrix) {

  return matrix;
}
```

# --solutions--

```js
function getNextLocation(matrix) {
  let prev = null;
  let curr = null;

  for (let row = 0; row < matrix.length; row++) {
    for (let col = 0; col < matrix[0].length; col++) {
      if (matrix[row][col] === 1) prev = [row, col];
      if (matrix[row][col] === 2) curr = [row, col];
    }
  }

  const [prevRow, prevCol] = prev;
  const [currRow, currCol] = curr;

  let dirX = currCol - prevCol;
  let dirY = currRow - prevRow;

  let nextRow = currRow + dirY;
  let nextCol = currCol + dirX;

  const maxRow = matrix.length - 1;
  const maxCol = matrix[0].length - 1;


  if (nextCol < 0 || nextCol > maxCol) dirX *= -1;
  if (nextRow < 0 || nextRow > maxRow) dirY *= -1;

  nextRow = currRow + dirY;
  nextCol = currCol + dirX;

  return [nextRow, nextCol];
}
```
