---
id: 69162d64f96574d9bb629efd
title: "Завдання 112: детектор штучного інтелекту"
challengeType: 28
dashedName: challenge-112
---

# --description--

Сьогоднішнє завдання натхненне випуском ChatGPT 30 листопада 2022 року.

Дано рядок з одного або кількох речень, визначте, чи ймовірно він був згенерований штучним інтелектом, використовуючи такі правила:

- Він містить два або більше дефіси (`-`).
- Він містить два або більше множини дужок (`()`). Текст може бути всередині дужок.
- Він містить три або більше слів з 7 або більше літерами.

- Слова розділені одним пробілом і складаються лише з літер (`A-Z`). Не включайте розділові знаки або інші нелітературні символи як частину слова.

Якщо задане речення відповідає будь-якому з наведених правил, поверніть `"AI"`, інакше поверніть `"Human"`.

# --hints--

`detectAI("The quick brown fox jumped over the lazy dog.")` має повертати `"Human"`.

```js
assert.equal(detectAI("The quick brown fox jumped over the lazy dog."), "Human");
```

`detectAI("The hypersonic brown fox - jumped (over) the lazy dog.")` має повертати `"Human"`.

```js
assert.equal(detectAI("The hypersonic brown fox - jumped (over) the lazy dog."), "Human");
```

`detectAI("Yes - you're right! I made a mistake there - let me try again.")` має повертати `"AI"`.

```js
assert.equal(detectAI("Yes - you're right! I made a mistake there - let me try again."), "AI");
```

`detectAI("The extraordinary students were studying vivaciously.")` має повертати `"AI"`.

```js
assert.equal(detectAI("The extraordinary students were studying vivaciously."), "AI");
```

`detectAI("The (excited) student was (coding) in the library.")` має повертати `"AI"`.

```js
assert.equal(detectAI("The (excited) student was (coding) in the library."), "AI");
```

# --seed--

## --seed-contents--

```js
function detectAI(text) {

  return text;
}
```

# --solutions--

```js
function detectAI(text) {
  const dashCount = (text.match(/-/g) || []).length;
  if (dashCount >= 2) return "AI";

  const parenCount = (text.match(/\([^)]*\)/g) || []).length;
  if (parenCount >= 2) return "AI";

  const words = text.split(" ");
  const longWordCount = words.filter(w => w.replace(/[^A-Za-z]/g, "").length >= 7).length;
  if (longWordCount >= 3) return "AI";

  return "Human";
}
```
