---
id: 69162d64f96574d9bb629efb
title: "Завдання 110: Відгадувач слів"
challengeType: 28
dashedName: challenge-110
---

# --description--

Дано два рядки однакової довжини: секретне слово та здогадку. Порівняйте здогадку із секретним словом за такими правилами:

- Секретне слово та здогадка складатимуться лише з літер верхнього регістру (`"A"` до `"Z"`);
- Для кожної літери у здогадці замініть її на число відповідно до того, як вона співпадає із секретним словом:
  - `"2"`, якщо літера є у секретному слові і знаходиться на правильній позиції.
  - `"1"`, якщо літера є у секретному слові, але на неправильній позиції.
  - `"0"`, якщо літера відсутня у секретному слові.
- Кожна літера у секретному слові може бути використана не більше одного разу.
- Спочатку призначаються точні співпадіння (`"2"`), потім часткові співпадіння (`"1"`) призначаються зліва направо для решти літер.
- Якщо літера зустрічається кілька разів у здогадці, вона може співпадати лише стільки разів, скільки разів вона є у секретному слові.

Наприклад, якщо секретне слово `"APPLE"`, а здогадка `"POPPA"`, поверніть `"10201"`:

Перша `"P"` не на правильній позиції (`"1"`), `"O"` відсутня у секретному слові (`"0"`), друга `"P"` на правильній позиції (`"2"`), третя `"P"` — це нуль (`"0"`), бо дві `"P"` у секретному слові вже використані, а `"A"` не на правильній позиції (`"1"`).

# --hints--

`compare("APPLE", "POPPA")` має повернути `"10201"`.

```js
assert.equal(compare("APPLE", "POPPA"), "10201");
```

`compare("REACT", "TRACE")` має повернути `"11221"`.

```js
assert.equal(compare("REACT", "TRACE"), "11221");
```

`compare("DEBUGS", "PYTHON")` має повернути `"000000"`.

```js
assert.equal(compare("DEBUGS", "PYTHON"), "000000");
```

`compare("JAVASCRIPT", "TYPESCRIPT")` має повернути `"0000222222"`.

```js
assert.equal(compare("JAVASCRIPT", "TYPESCRIPT"), "0000222222");
```

`compare("ORANGE", "ROUNDS")` має повернути `"110200"`.

```js
assert.equal(compare("ORANGE", "ROUNDS"), "110200");
```

`compare("WIRELESS", "ETHERNET")` має повернути `"10021000"`.

```js
assert.equal(compare("WIRELESS", "ETHERNET"), "10021000");
```

# --seed--

## --seed-contents--

```js
function compare(word, guess) {

  return word;
}
```

# --solutions--

```js
function compare(word, guess) {
  const n = word.length;
  let result = Array(n).fill("0");
  let used = Array(n).fill(false);

  for (let i = 0; i < n; i++) {
    if (guess[i] === word[i]) {
      result[i] = "2";
      used[i] = true;
    }
  }

  for (let i = 0; i < n; i++) {
    if (result[i] === "0") {
      for (let j = 0; j < n; j++) {
        if (!used[j] && guess[i] === word[j]) {
          result[i] = "1";
          used[j] = true;
          break;
        }
      }
    }
  }

  return result.join("");
}
```
