---
id: 69162d64f96574d9bb629efe
title: "Завдання 113: милі в кілометри"
challengeType: 28
dashedName: challenge-113
---

# --description--

Дано відстань у милях як число, поверніть еквівалентну відстань у кілометрах.

- Вхідні дані завжди будуть невід’ємним числом.
- 1 миля дорівнює 1.60934 кілометра.
- Округліть результат до двох десяткових знаків.
- Видаліть непотрібні нулі в кінці округленого результату.

# --hints--

`convertToKm(1)` має повертати `1.61`.

```js
assert.equal(convertToKm(1), 1.61);
```

`convertToKm(21)` має повертати `33.8`.

```js
assert.equal(convertToKm(21), 33.8);
```

`convertToKm(3.5)` має повертати `5.63`.

```js
assert.equal(convertToKm(3.5), 5.63);
```

`convertToKm(0)` має повертати `0`.

```js
assert.equal(convertToKm(0), 0);
```

`convertToKm(0.621371)` має повертати `1`.

```js
assert.equal(convertToKm(0.621371), 1);
```

# --seed--

## --seed-contents--

```js
function convertToKm(miles) {

  return miles;
}
```

# --solutions--

```js
function convertToKm(miles) {
  const km = miles * 1.60934;
  return Math.round(km * 100) / 100;
}
```
