---
id: 69162d64f96574d9bb629f01
title: "Завдання 116: підрахунок перестановок"
challengeType: 28
dashedName: challenge-116
---

# --description--

Дано рядок, поверніть число унікальних перестановок, які можна утворити з його символів.

- Перестановка — це будь-який перерозподіл символів у рядку.
- Не рахуйте дублікати перестановок.
- Якщо рядок містить повторювані символи, повторювані варіанти слід враховувати лише один раз.
- Рядок міститиме лише літери (`A-Z`, `a-z`).

Наприклад, для `"abb"` поверніть `3`, бо існує три унікальні способи впорядкувати літери: `"abb"`, `"bab"` та `"bba"`.

# --hints--

`countPermutations("abb")` має повернути `3`.

```js
assert.equal(countPermutations("abb"), 3);
```

`countPermutations("abc")` має повернути `6`.

```js
assert.equal(countPermutations("abc"), 6);
```

`countPermutations("racecar")` має повернути `630`.

```js
assert.equal(countPermutations("racecar"), 630);
```

`countPermutations("freecodecamp")` має повернути `39916800`.

```js
assert.equal(countPermutations("freecodecamp"), 39916800);
```

# --seed--

## --seed-contents--

```js
function countPermutations(str) {

  return str;
}
```

# --solutions--

```js
function countPermutations(str) {
  const freq = {};
  for (const ch of str) {
    freq[ch] = (freq[ch] || 0) + 1;
  }

  function factorial(n) {
    let result = 1;
    for (let i = 2; i <= n; i++) result *= i;
    return result;
  }

  const n = str.length;
  let result = factorial(n);

  for (const ch in freq) {
    result /= factorial(freq[ch]);
  }

  return result;
}
```
