---
id: 69162d64f96574d9bb629f00
title: "Завдання 115: конвертер елемента впорядкованого списку Markdown"
challengeType: 28
dashedName: challenge-115
---

# --description--

Дано рядок, що представляє елемент впорядкованого списку в Markdown. Потрібно повернути еквівалентний HTML-рядок.

Дійсний елемент впорядкованого списку в Markdown повинен:

- Починатися з нуля або більше пробілів, за якими йде
- Число (1 або більше) і крапка (`.`), після чого
- Принаймні один пробіл, а потім
- Текст елемента списку.

Якщо рядок не має точно такого формату, поверніть `"Invalid format"`. Інакше обгорніть текст елемента списку в теги `li` і поверніть рядок.

Наприклад, для `"1. My item"` поверніть `"<li>My item</li>"`.

Примітка: консоль може не відображати HTML-теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`convertListItem("1. My item")` має повернути `"<li>My item</li>"`.

```js
assert.equal(convertListItem("1. My item"), "<li>My item</li>");
```

`convertListItem(" 1.  Another item")` має повернути `"<li>Another item</li>"`.

```js
assert.equal(convertListItem(" 1.  Another item"), "<li>Another item</li>");
```

`convertListItem("1 . invalid item")` має повернути `"Invalid format"`.

```js
assert.equal(convertListItem("1 . invalid item"), "Invalid format");
```

`convertListItem("2. list item text")` має повернути `"<li>list item text</li>"`.

```js
assert.equal(convertListItem("2. list item text"), "<li>list item text</li>");
```

`convertListItem(". invalid again")` має повернути `"Invalid format"`.

```js
assert.equal(convertListItem(". invalid again"), "Invalid format");
```

`convertListItem("A. last invalid")` має повернути `"Invalid format"`.

```js
assert.equal(convertListItem("A. last invalid"), "Invalid format");
```

# --seed--

## --seed-contents--

```js
function convertListItem(markdown) {

  return markdown;
}
```

# --solutions--

```js
function convertListItem(markdown) {
  const match = markdown.match(/^\s*(\d+)\.\s+(.+)$/);
  if (!match) return "Invalid format";

  const text = match[2];
  return `<li>${text}</li>`;
}
```
