---
id: 69162d64f96574d9bb629f04
title: "Завдання 119: стиснення рядка"
challengeType: 28
dashedName: challenge-119
---

# --description--

Дано рядок речення, поверніть стиснену версію речення, де послідовні повторювані слова замінені словом, за яким у круглих дужках вказано кількість повторів.

- Стискаються лише послідовні дублікати.
- Слова розділені одним пробілом.

Наприклад, для `"yes yes yes please"` поверніть `"yes(3) please"`.

# --hints--

`compressString("yes yes yes please")` має повернути `"yes(3) please"`.

```js
assert.equal(compressString("yes yes yes please"), "yes(3) please");
```

`compressString("I have have have apples")` має повернути `"I have(3) apples"`.

```js
assert.equal(compressString("I have have have apples"), "I have(3) apples");
```

`compressString("one one three and to the the the the")` має повернути `"one(2) three and to the(4)"`.

```js
assert.equal(compressString("one one three and to the the the the"), "one(2) three and to the(4)");
```

`compressString("route route route route route route tee tee tee tee tee tee")` має повернути `"route(6) tee(6)"`.

```js
assert.equal(compressString("route route route route route route tee tee tee tee tee tee"), "route(6) tee(6)");
```

# --seed--

## --seed-contents--

```js
function compressString(sentence) {

  return sentence;
}
```

# --solutions--

```js
function compressString(sentence) {
  const words = sentence.split(" ");
  const result = [];
  let count = 1;

  for (let i = 0; i < words.length; i++) {
    if (words[i] === words[i + 1]) {
      count++;
    } else {
      if (count > 1) {
        result.push(`${words[i]}(${count})`);
      } else {
        result.push(words[i]);
      }
      count = 1;
    }
  }

  return result.join(" ");
}
```
