---
id: 69162d64f96574d9bb629eff
title: "Завдання 114: camel case у snake case"
challengeType: 28
dashedName: challenge-114
---

# --description--

Дано рядок у camel case. Поверніть версію цього рядка у snake case, використовуючи такі правила:

- Вхідний рядок міститиме лише літери (`A-Z` та `a-z`) і завжди починатиметься з літери нижнього регістру.
- Кожна велика літера в camel case позначає початок нового слова.
- Перетворіть усі літери на нижній регістр.
- Розділяйте слова символом підкреслення (`_`).

# --hints--

`toSnake("helloWorld")` має повернути `"hello_world"`.

```js
assert.equal(toSnake("helloWorld"), "hello_world");
```

`toSnake("myVariableName")` має повернути `"my_variable_name"`.

```js
assert.equal(toSnake("myVariableName"), "my_variable_name");
```

`toSnake("freecodecampDailyChallenges")` має повернути `"freecodecamp_daily_challenges"`.

```js
assert.equal(toSnake("freecodecampDailyChallenges"), "freecodecamp_daily_challenges");
```

# --seed--

## --seed-contents--

```js
function toSnake(camel) {

  return camel;
}
```

# --solutions--

```js
function toSnake(camel) {
  return camel
    .replace(/([A-Z])/g, "_$1")
    .toLowerCase();
}
```
