---
id: 69162d64f96574d9bb629f02
title: "Завдання 117: симетрична різниця"
challengeType: 28
dashedName: challenge-117
---

# --description--

Дано два масиви, поверніть новий масив, що містить їхню симетричну різницю.

- Симетрична різниця між двома множинами — це множина значень, які з’являються в одній із множин, але не в обох одночасно.
- Поверніть значення в тому порядку, в якому вони вперше з’являються у вхідних масивах.

# --hints--

`difference([1, 2, 3], [3, 4, 5])` має повернути `[1, 2, 4, 5]`.

```js
assert.deepEqual(difference([1, 2, 3], [3, 4, 5]), [1, 2, 4, 5]);
```

`difference(["a", "b"], ["c", "b"])` має повернути `["a", "c"]`.

```js
assert.deepEqual(difference(["a", "b"], ["c", "b"]), ["a", "c"]);
```

`difference([1, "a", 2], [2, "b", "a"])` має повернути `[1, "b"]`.

```js
assert.deepEqual(difference([1, "a", 2], [2, "b", "a"]), [1, "b"]);
```

`difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9])` має повернути `[2, 4, 6, 8]`.

```js
assert.deepEqual(difference([1, 3, 5, 7, 9], [1, 2, 3, 4, 5, 6, 7, 8, 9]), [2, 4, 6, 8]);
```

# --seed--

## --seed-contents--

```js
function difference(arr1, arr2) {

  return arr1;
}
```

# --solutions--

```js
function difference(arr1, arr2) {
  const diff = [];

  for (const v of arr1) {

    if (!diff.includes(v) && !arr2.includes(v)) diff.push(v)
  }
  
  for (const v of arr2) {
    if (!diff.includes(v) && !arr1.includes(v)) diff.push(v)
  }

  return diff;
}
```
