---
id: 691b559495c5cb5a37b9b485
title: "Завдання 125: гра життя"
challengeType: 28
dashedName: challenge-125
---

# --description--

Дано матрицю (масив масивів), що представляє поточний стан у грі життя Конвея. Поверніть наступний стан матриці, використовуючи такі правила:

- Кожна клітинка або `1` (жива), або `0` (мертва).
- Сусідами клітинки є до восьми навколишніх клітинок (вертикально, горизонтально та по діагоналі).
- Клітинки на краях мають менше ніж вісім сусідів.

Правила оновлення кожної клітинки:

- Будь-яка жива клітинка з менше ніж двома живими сусідами вмирає (недостатня кількість).
- Будь-яка жива клітинка з двома або трьома живими сусідами продовжує жити.
- Будь-яка жива клітинка з більш ніж трьома живими сусідами вмирає (надмірна кількість).
- Будь-яка мертва клітинка з рівно трьома живими сусідами стає живою (розмноження).

Наприклад, якщо задано:

```json
[
  [0, 1, 0],
  [0, 1, 1],
  [1, 1, 0]
]
```

поверніть:

```json
[
  [0, 1, 1],
  [0, 0, 1],
  [1, 1, 1]
]
```

Кожна клітинка оновлюється відповідно до кількості живих сусідів. Наприклад, `[0][0]` залишається мертвою (2 живих сусіди), `[0][1]` залишається живою (2 живих сусіди), `[0][2]` вмирає (3 живих сусіди) і так далі.

# --hints--

`gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]])` має повернути `[[0, 1, 1], [0, 0, 1], [1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 0], [0, 1, 1], [1, 1, 0]]), [[0, 1, 1], [0, 0, 1], [1, 1, 1]]);
```

`gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]])` має повернути `[[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]`.

```js
assert.deepEqual(gameOfLife([[1, 1, 0, 0], [1, 0, 1, 0], [0, 1, 1, 1], [0, 0, 1, 0]]), [[1, 1, 0, 0], [1, 0, 0, 1], [0, 0, 0, 1], [0, 1, 1, 1]]);
```

`gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]])` має повернути `[[0, 0, 0], [0, 1, 0], [0, 0, 0]]`.

```js
assert.deepEqual(gameOfLife([[1, 0, 0], [0, 1, 0], [0, 0, 1]]), [[0, 0, 0], [0, 1, 0], [0, 0, 0]]);
```

`gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]])` має повернути `[[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]`.

```js
assert.deepEqual(gameOfLife([[0, 1, 1, 0], [1, 1, 0, 1], [0, 1, 1, 0], [0, 0, 1, 0]]), [[1, 1, 1, 0], [1, 0, 0, 1], [1, 0, 0, 1], [0, 1, 1, 0]]);
```

# --seed--

## --seed-contents--

```js
function gameOfLife(grid) {
  return grid;
}
```

# --solutions--

```js
function gameOfLife(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const countLiveNeighbors = (r, c) => {
    let count = 0;
    for (let i = r - 1; i <= r + 1; i++) {
      for (let j = c - 1; j <= c + 1; j++) {
        if (
          i >= 0 &&
          i < rows &&
          j >= 0 &&
          j < cols &&
          !(i === r && j === c) &&
          grid[i][j] === 1
        ) {
          count++;
        }
      }
    }
    return count;
  };

  const next = grid.map(row => [...row]);

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      const liveNeighbors = countLiveNeighbors(r, c);

      if (grid[r][c] === 1) {
        if (liveNeighbors < 2 || liveNeighbors > 3) next[r][c] = 0;
      } else {
        if (liveNeighbors === 3) next[r][c] = 1;
      }
    }
  }

  return next;
}
```
