---
id: 691b559495c5cb5a37b9b487
title: "Завдання 127: перевірка швидкості"
challengeType: 28
dashedName: challenge-127
---

# --description--

Враховуючи швидкість, з якою ви рухаєтеся в милях на годину (MPH), та обмеження швидкості в кілометрах на годину (KPH), визначте, чи перевищуєте ви швидкість, і чи отримаєте попередження або штраф.

- 1 миля дорівнює 1,60934 кілометра.
- Якщо ви рухаєтеся зі швидкістю, меншою або рівною обмеженню, поверніть `"Not Speeding"`.
- Якщо ви рухаєтеся зі швидкістю, що перевищує обмеження на 5 KPH або менше, поверніть `"Warning"`.
- Якщо ви рухаєтеся зі швидкістю, що перевищує обмеження більш ніж на 5 KPH, поверніть `"Ticket"`.

# --hints--

`speedCheck(30, 70)` має повертати `"Not Speeding"`.

```js
assert.equal(speedCheck(30, 70), "Not Speeding");
```

`speedCheck(40, 60)` має повертати `"Warning"`.

```js
assert.equal(speedCheck(40, 60), "Warning");
```

`speedCheck(40, 65)` має повертати `"Not Speeding"`.

```js
assert.equal(speedCheck(40, 65), "Not Speeding");
```

`speedCheck(60, 90)` має повертати `"Ticket"`.

```js
assert.equal(speedCheck(60, 90), "Ticket");
```

`speedCheck(65, 100)` має повертати `"Warning"`.

```js
assert.equal(speedCheck(65, 100), "Warning");
```

`speedCheck(88, 40)` має повертати `"Ticket"`.

```js
assert.equal(speedCheck(88, 40), "Ticket");
```

# --seed--

## --seed-contents--

```js
function speedCheck(speedMph, speedLimitKph) {
  return speedMph;
}
```

# --solutions--

```js
function speedCheck(speedMph, speedLimitKph) {
  const speedKph = speedMph * 1.60934;

  if (speedKph <= speedLimitKph) {
    return "Not Speeding";
  }

  const over = speedKph - speedLimitKph;

  if (over <= 5) {
    return "Warning";
  }

  return "Ticket";
}
```
