---
id: 691b559495c5cb5a37b9b480
title: "Завдання 120: фунти в кілограми"
challengeType: 28
dashedName: challenge-120
---

# --description--

Дано вагу в фунтах у вигляді числа, поверніть рядок `"(lbs) pounds equals (kgs) kilograms."`.

- Замініть `"(lbs)"` на вхідне число.
- Замініть `"(kgs)"` на вхідне число, конвертоване в кілограми, округлене до двох десяткових знаків і завжди з двома десятковими знаками у значенні.
- 1 фунт дорівнює 0.453592 кілограма.
- Якщо вхідні дані — це `1`, використовуйте `"pound"` замість `"pounds"`.
- Якщо конвертоване значення — це `1`, використовуйте `"kilogram"` замість `"kilograms"`.

# --hints--

`convertToKgs(1)` має повертати `"1 pound equals 0.45 kilograms."`.

```js
assert.equal(convertToKgs(1), "1 pound equals 0.45 kilograms.");
```

`convertToKgs(0)` має повертати `"0 pounds equals 0.00 kilograms."`.

```js
assert.equal(convertToKgs(0), "0 pounds equals 0.00 kilograms.");
```

`convertToKgs(100)` має повертати `"100 pounds equals 45.36 kilograms."`.

```js
assert.equal(convertToKgs(100), "100 pounds equals 45.36 kilograms.");
```

`convertToKgs(2.5)` має повертати `"2.5 pounds equals 1.13 kilograms."`.

```js
assert.equal(convertToKgs(2.5), "2.5 pounds equals 1.13 kilograms.");
```

`convertToKgs(2.20462)` має повертати `"2.20462 pounds equals 1.00 kilogram."`.

```js
assert.equal(convertToKgs(2.20462), "2.20462 pounds equals 1.00 kilogram.");
```

# --seed--

## --seed-contents--

```js
function convertToKgs(lbs) {
  return lbs;
}
```

# --solutions--

```js
function convertToKgs(lbs) {
  const KG_PER_POUND = 0.453592;
  const kgs = (lbs * KG_PER_POUND).toFixed(2);

  const poundWord = lbs === 1 ? "pound" : "pounds";
  const kilogramWord = kgs === "1.00" ? "kilogram" : "kilograms";

  return `${lbs} ${poundWord} equals ${kgs} ${kilogramWord}.`;
}
```
