---
id: 691b559495c5cb5a37b9b483
title: "Завдання 123: конструктор римських чисел"
challengeType: 28
dashedName: challenge-123
---

# --description--

Дано ціле число, поверніть його еквівалентне значення римськими цифрами.

Римські цифри використовують такі символи:

| Символ | Значення |
| ------ | -------- |
| I      | 1        |
| V      | 5        |
| X      | 10       |
| L      | 50       |
| C      | 100      |
| D      | 500      |
| M      | 1000     |

Римські цифри записують від найбільшого до найменшого, зліва направо, за такими правилами:

- Додавання використовується, коли символ слідує за символом з таким самим або меншим значенням. Наприклад, `18` записується як `XVIII` (10 + 5 + 1 + 1 + 1 = 18).
- Віднімання використовується, коли менший символ стоїть перед більшим, щоб представити 4 або 9 у будь-якому розряді. Наприклад, 19 записується як `XIX` (10 + (10 - 1)).
- Жоден символ не може повторюватися більше трьох разів підряд. Віднімання використовується, коли інакше довелося б записати символ більше трьох разів підряд. Тож найбільше число, яке можна записати, — 3999.

Ось ще один приклад: для `1464` поверніть `"MCDLXIV"` (1000 + (500 - 100) + 50 + 10 + (5 - 1)).

# --hints--

`toRoman(18)` має повернути `"XVIII"`.

```js
assert.equal(toRoman(18), "XVIII");
```

`toRoman(19)` має повернути `"XIX"`.

```js
assert.equal(toRoman(19), "XIX");
```

`toRoman(1464)` має повернути `"MCDLXIV"`.

```js
assert.equal(toRoman(1464), "MCDLXIV");
```

`toRoman(2025)` має повернути `"MMXXV"`.

```js
assert.equal(toRoman(2025), "MMXXV");
```

`toRoman(3999)` має повернути `"MMMCMXCIX"`.

```js
assert.equal(toRoman(3999), "MMMCMXCIX");
```

# --seed--

## --seed-contents--

```js
function toRoman(num) {
  return num;
}
```

# --solutions--

```js
function toRoman(num) {
  const symbols = [
    ["M", 1000],
    ["CM", 900],
    ["D", 500],
    ["CD", 400],
    ["C", 100],
    ["XC", 90],
    ["L", 50],
    ["XL", 40],
    ["X", 10],
    ["IX", 9],
    ["V", 5],
    ["IV", 4],
    ["I", 1]
  ];

  let result = "";
  for (const [sym, val] of symbols) {
    while (num >= val) {
      result += sym;
      num -= val;
    }
  }
  return result;
}
```
