---
id: 691b559495c5cb5a37b9b488
title: "Завдання 128: підрахунок приголосних"
challengeType: 28
dashedName: challenge-128
---

# --description--

Дано рядок і цільове число. Визначте, чи містить рядок рівно цільове число приголосних.

- Приголосні — це всі алфавітні символи, крім `"a"`, `"e"`, `"i"`, `"o"` та `"u"` у будь-якому регістрі.
- Ігноруйте цифри, розділові знаки, пробіли та інші не буквені символи під час підрахунку.

# --hints--

`hasConsonantCount("helloworld", 7)` має повернути `true`.

```js
assert.isTrue(hasConsonantCount("helloworld", 7));
```

`hasConsonantCount("eieio", 5)` має повернути `false`.

```js
assert.isFalse(hasConsonantCount("eieio", 5));
```

`hasConsonantCount("freeCodeCamp Rocks!", 11)` має повернути `true`.

```js
assert.isTrue(hasConsonantCount("freeCodeCamp Rocks!", 11));
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)` має повернути `false`.

```js
assert.isFalse(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 24)
);
```

`hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)` має повернути `true`.

```js
assert.isTrue(
  hasConsonantCount("Th3 Qu!ck Br0wn F0x Jump5 0ver Th3 L@zy D0g.", 23)
);
```

# --seed--

## --seed-contents--

```js
function hasConsonantCount(text, target) {
  return text;
}
```

# --solutions--

```js
function hasConsonantCount(text, target) {
  const vowels = new Set(["a", "e", "i", "o", "u"]);
  let count = 0;

  for (const char of text.toLowerCase()) {
    if (/[a-z]/.test(char) && !vowels.has(char)) {
      count++;
    }
  }

  return count === target;
}
```
