---
id: 691b559495c5cb5a37b9b482
title: "Завдання 122: парсер жирного тексту в markdown"
challengeType: 28
dashedName: challenge-122
---

# --description--

Дано рядок, який може містити жирний текст у markdown. Поверніть еквівалентний HTML-рядок.

- Жирний текст у markdown — це будь-який текст, який починається і закінчується двома зірочками (`**`) або двома підкресленнями (`__`).
- Між текстом і зірочками або підкресленнями не може бути пробілів, але пробіли можуть бути всередині самого тексту.
- Перетворіть усі випадки жирного тексту на HTML-теги `b` і поверніть рядок.

Наприклад, для `"**This is bold**"` поверніть `"<b>This is bold</b>"`.

Примітка: консоль може не відображати HTML-теги в рядках під час виводу повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseBold("**This is bold**")` має повернути `"<b>This is bold</b>"`.

```js
assert.equal(parseBold("**This is bold**"), "<b>This is bold</b>");
```

`parseBold("__This is also bold__")` має повернути `"<b>This is also bold</b>"`.

```js
assert.equal(parseBold("__This is also bold__"), "<b>This is also bold</b>");
```

`parseBold("**This is not bold **")` має повернути `"**This is not bold **"`.

```js
assert.equal(parseBold("**This is not bold **"), "**This is not bold **");
```

`parseBold("__ This is also not bold__")` має повернути `"__ This is also not bold__"`.

```js
assert.equal(parseBold("__ This is also not bold__"), "__ This is also not bold__");
```

`parseBold("The **quick** brown fox __jumps__ over the **lazy** dog.")` має повернути `"The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog."`.

```js
assert.equal(parseBold("The **quick** brown fox __jumps__ over the **lazy** dog."), "The <b>quick</b> brown fox <b>jumps</b> over the <b>lazy</b> dog.");
```

# --seed--

## --seed-contents--

```js
function parseBold(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBold(markdown) {
  markdown = markdown.replace(/\*\*(\S(?:.*?\S)?)\*\*/g, "<b>$1</b>");
  markdown = markdown.replace(/__(\S(?:.*?\S)?)__/g, "<b>$1</b>");

  return markdown;
}
```
