---
id: 691b559495c5cb5a37b9b489
title: "Завдання 129: парсер блокових цитат у markdown"
challengeType: 28
dashedName: challenge-129
---

# --description--

Дано рядок, який містить блокову цитату у markdown. Потрібно повернути еквівалентний HTML-рядок.

Блокова цитата у markdown — це будь-який рядок, який:

- Починається з нуля або більше пробілів
- Далі йде знак більше (`>`)
- Потім один або більше пробілів
- І нарешті текст блокової цитати.

Поверніть текст блокової цитати, оточений відкриваючим і закриваючим HTML-тегом `blockquote`.

Наприклад, для `"> This is a quote"` поверніть `<blockquote>This is a quote</blockquote>`.

Примітка: консоль може не відображати HTML-теги у рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseBlockquote("> This is a quote")` має повернути `"<blockquote>This is a quote</blockquote>"`.

```js
assert.equal(parseBlockquote("> This is a quote"), "<blockquote>This is a quote</blockquote>");
```

`parseBlockquote(" > This is also a quote")` має повернути `"<blockquote>This is also a quote</blockquote>"`.

```js
assert.equal(parseBlockquote(" > This is also a quote"), "<blockquote>This is also a quote</blockquote>");
```

`parseBlockquote("  >    So  Is  This")` має повернути `"<blockquote>So  Is  This</blockquote>"`.

```js
assert.equal(parseBlockquote("  >    So  Is  This"), "<blockquote>So  Is  This</blockquote>");
```

# --seed--

## --seed-contents--

```js
function parseBlockquote(markdown) {
  return markdown;
}
```

# --solutions--

```js
function parseBlockquote(markdown) {
  const blockquoteRegex = /^\s*>\s+(.+)$/;
  const match = markdown.match(blockquoteRegex);

  const text = match[1];
  return `<blockquote>${text}</blockquote>`;
}
```
