---
id: 691b559495c5cb5a37b9b486
title: "Завдання 126: зробити першу літеру великою"
challengeType: 28
dashedName: challenge-126
---

# --description--

Дано рядок title, поверніть новий рядок, відформатований у форматі заголовка за такими правилами:

- Зробіть першу літеру кожного слова великою.
- Всі інші літери в кожному слові зробіть у нижньому регістрі.
- Слова завжди розділені одним пробілом.

# --hints--

`titleCase("hello world")` має повернути `"Hello World"`.

```js
assert.equal(titleCase("hello world"), "Hello World");
```

`titleCase("the quick brown fox")` має повернути `"The Quick Brown Fox"`.

```js
assert.equal(titleCase("the quick brown fox"), "The Quick Brown Fox");
```

`titleCase("JAVASCRIPT AND PYTHON")` має повернути `"Javascript And Python"`.

```js
assert.equal(titleCase("JAVASCRIPT AND PYTHON"), "Javascript And Python");
```

`titleCase("AvOcAdO tOAst fOr brEAkfAst")` має повернути `"Avocado Toast For Breakfast"`.

```js
assert.equal(titleCase("AvOcAdO tOAst fOr brEAkfAst"), "Avocado Toast For Breakfast");
```

# --seed--

## --seed-contents--

```js
function titleCase(title) {
  return title;
}
```

# --solutions--

```js
function titleCase(title) {
  return title
    .split(" ")
    .map(word => word.charAt(0).toUpperCase() + word.slice(1).toLowerCase())
    .join(" ");
}
```
