---
id: 6925e2068081f40f549ced1b
title: "Завдання 137: генератор сніжинок"
challengeType: 28
dashedName: challenge-137
---

# --description--

Дано багаторядковий рядок, який використовує символи нового рядка (`\n`) для позначення розриву рядка. Поверніть новий рядок, у якому кожен рядок віддзеркалений по горизонталі та приєднаний до кінця оригінального рядка.

- Віддзеркальте рядок, перевернувши всі його символи, включно з пробілами.

Наприклад, для `"* \n *\n* "`, який виводить у консоль:

```sh
* 
 *
* 
```

Поверніть `"*  *\n ** \n*  *"`, який виводить у консоль:

```sh
*  *
 ** 
*  *
```

Уважно зверніть увагу на пробіли у заданих і повернутих рядках. Обов’язково не обрізайте жоден із них.

# --hints--

`generateSnowflake("* \n *\n* ")` має повернути `"*  *\n ** \n*  *"`.

```js
assert.equal(generateSnowflake("* \n *\n* "), "*  *\n ** \n*  *");
```

`generateSnowflake("X=~")` має повернути `"X=~~=X"`.

```js
assert.equal(generateSnowflake("X=~"), "X=~~=X");
```

`generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  ")` має повернути `" X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X "`.

```js
assert.equal(generateSnowflake(" X  \n  v \nX--=\n  ^ \n X  "), " X    X \n  v  v  \nX--==--X\n  ^  ^  \n X    X ");
```

`generateSnowflake("*   *\n * * \n* * *\n * * \n*   *")` має повернути `"*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *"`.

```js
assert.equal(generateSnowflake("*   *\n * * \n* * *\n * * \n*   *"), "*   **   *\n * *  * * \n* * ** * *\n * *  * * \n*   **   *");
```

`generateSnowflake("*  -\n * -\n*  -")` має повернути `"*  --  *\n * -- * \n*  --  *"`.

```js
assert.equal(generateSnowflake("*  -\n * -\n*  -"), "*  --  *\n * -- * \n*  --  *");
```

# --seed--

## --seed-contents--

```js
function generateSnowflake(crystals) {

  return crystals;
}
```

# --solutions--

```js
function generateSnowflake(crystals) {
  const lines = crystals.split("\n");

  const mirroredLines = lines.map(line => {
    const reversed = line.split("").reverse().join("");
    return line + reversed;
  });

  return mirroredLines.join("\n");
}
```
