---
id: 6925e2068081f40f549ced1c
title: "Завдання 138: сума дільників"
challengeType: 28
dashedName: challenge-138
---

# --description--

Дано додатне ціле число, поверніть суму всіх його дільників.

- Дільник — це будь-яке ціле число, яке ділить задане число без остачі (остача дорівнює `0`).
- Рахуйте кожен дільник лише один раз.

Наприклад, для `6` поверніть `12`, тому що дільниками `6` є `1`, `2`, `3` і `6`, а їх сума дорівнює `12`.

# --hints--

`sumDivisors(6)` має повернути `12`.

```js
assert.equal(sumDivisors(6), 12);
```

`sumDivisors(13)` має повернути `14`.

```js
assert.equal(sumDivisors(13), 14);
```

`sumDivisors(28)` має повернути `56`.

```js
assert.equal(sumDivisors(28), 56);
```

`sumDivisors(84)` має повернути `224`.

```js
assert.equal(sumDivisors(84), 224);
```

`sumDivisors(549)` має повернути `806`.

```js
assert.equal(sumDivisors(549), 806);
```

`sumDivisors(9348)` має повернути `23520`.

```js
assert.equal(sumDivisors(9348), 23520);
```

# --seed--

## --seed-contents--

```js
function sumDivisors(n) {

  return n;
}
```

# --solutions--

```js
function sumDivisors(n) {
  let sum = 0;
  for (let i = 1; i <= n; i++) {
    if (n % i === 0) {
      sum += i;
    }
  }
  return sum;
}
```
