---
id: 6925e2068081f40f549ced1a
title: "Завдання 136: парсер зображень у markdown"
challengeType: 28
dashedName: challenge-136
---

# --description--

Дано рядок зображення у форматі Markdown, поверніть еквівалентний HTML-рядок.

Зображення в Markdown має такий формат: `"![alt text](image_url)"`. Де:

- `alt text` — це опис зображення (значення атрибута `alt`).
- `image_url` — це URL-адреса джерела зображення (значення атрибута `src`).

Поверніть рядок HTML-тега `img` з множиною атрибутів, де `src` встановлено на URL зображення, а `alt` — на альтернативний текст.

Наприклад, для `"![Cute cat](cat.png)"` поверніть `'<img src="cat.png" alt="Cute cat">'`;

- Переконайтеся, що тег, порядок атрибутів, пробіли та використання лапок збігаються з наведеним вище.

Примітка: консоль може не відображати HTML-теги в рядках під час виведення повідомлень — перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseImage("![Cute cat](cat.png)")` має повертати `'<img src="cat.png" alt="Cute cat">'`.

```js
assert.equal(parseImage("![Cute cat](cat.png)"), '<img src="cat.png" alt="Cute cat">');
```

`parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)")` має повертати `'<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">'`.

```js
assert.equal(parseImage("![Rocket Ship](https://freecodecamp.org/cdn/rocket-ship.jpg)"), '<img src="https://freecodecamp.org/cdn/rocket-ship.jpg" alt="Rocket Ship">');
```

`parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)")` має повертати `'<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">'`.

```js
assert.equal(parseImage("![Cute cats!](https://freecodecamp.org/cats.jpeg)"), '<img src="https://freecodecamp.org/cats.jpeg" alt="Cute cats!">');
```

# --seed--

## --seed-contents--

```js
function parseImage(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseImage(markdown) {
  const regex = /!\[(.*?)\]\((.*?)\)/;
  const match = markdown.match(regex);
  if (!match) return "";
  const alt = match[1];
  const src = match[2];
  return `<img src="${src}" alt="${alt}">`;
}
```
