---
id: 6925e2068081f40f549ced1d
title: "Завдання 139: камінь, ножиці, папір"
challengeType: 28
dashedName: challenge-139
---

# --description--

Дано два рядки: перший представляє гравця 1, а другий — гравця 2. Визначте переможця у грі камінь, ножиці, папір.

- Вхідні рядки завжди будуть `"Rock"`, `"Paper"` або `"Scissors"`.
- `"Rock"` перемагає `"Scissors"`.
- `"Paper"` перемагає `"Rock"`.
- `"Scissors"` перемагає `"Paper"`.

Поверніть:

- `"Player 1 wins"`, якщо виграє гравець 1.
- `"Player 2 wins"`, якщо виграє гравець 2.
- `"Tie"`, якщо обидва гравці вибрали однаковий варіант.

# --hints--

`rockPaperScissors("Rock", "Rock")` має повернути `"Tie"`.

```js
assert.equal(rockPaperScissors("Rock", "Rock"), "Tie");
```

`rockPaperScissors("Rock", "Paper")` має повернути `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Paper"), "Player 2 wins");
```

`rockPaperScissors("Scissors", "Paper")` має повернути `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Paper"), "Player 1 wins");
```

`rockPaperScissors("Rock", "Scissors")` має повернути `"Player 1 wins"`.

```js
assert.equal(rockPaperScissors("Rock", "Scissors"), "Player 1 wins");
```

`rockPaperScissors("Scissors", "Scissors")` має повернути `"Tie"`.

```js
assert.equal(rockPaperScissors("Scissors", "Scissors"), "Tie");
```

`rockPaperScissors("Scissors", "Rock")` має повернути `"Player 2 wins"`.

```js
assert.equal(rockPaperScissors("Scissors", "Rock"), "Player 2 wins");
```

# --seed--

## --seed-contents--

```js
function rockPaperScissors(player1, player2) {

  return player1;
}
```

# --solutions--

```js
function rockPaperScissors(player1, player2) {
  if (player1 === player2) return "Tie";

  if (
    (player1 === "Rock" && player2 === "Scissors") ||
    (player1 === "Paper" && player2 === "Rock") ||
    (player1 === "Scissors" && player2 === "Paper")
  ) {
    return "Player 1 wins";
  } else {
    return "Player 2 wins";
  }
}
```
