---
id: 69272dcf1c24b44fd79137c5
title: "Завдання 142: Сума рядка"
challengeType: 28
dashedName: challenge-142
---

# --description--

Дано рядок, що містить цифри та інші символи. Поверніть суму всіх чисел у рядку.

- Розглядайте послідовні цифри як одне число. Наприклад, `"13"` рахується як 13, а не 1 + 3.
- Ігноруйте всі символи, які не є цифрами.

# --hints--

`stringSum("3apples2bananas")` має повернути `5`.

```js
assert.equal(stringSum("3apples2bananas"), 5);
```

`stringSum("10cats5dogs2birds")` має повернути `17`.

```js
assert.equal(stringSum("10cats5dogs2birds"), 17);
```

`stringSum("125344")` має повернути `125344`.

```js
assert.equal(stringSum("125344"), 125344);
```

`stringSum("a1b20c300")` має повернути `321`.

```js
assert.equal(stringSum("a1b20c300"), 321);
```

`stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5")` має повернути `1653`.

```js
assert.equal(stringSum("a12b34c56d78e90f123g456h789i0j1k2l3m4n5"), 1653);
```

# --seed--

## --seed-contents--

```js
function stringSum(str) {

  return str;
}
```

# --solutions--

```js
function stringSum(str) {
  const matches = str.match(/\d+/g); // find all consecutive digits
  return matches.reduce((sum, num) => sum + Number(num), 0);
}
```
