---
id: 69272dcf1c24b44fd79137c4
title: "Завдання 141: пальне для зльоту"
challengeType: 28
dashedName: challenge-141
---

# --description--

Враховуючи кількість галонів пального, що наразі є у вашому літаку, та необхідну кількість літрів пального для досягнення пункту призначення, визначте, скільки додаткових галонів пального слід додати.

- 1 галон дорівнює 3.78541 літрів.
- Якщо в літаку вже достатньо пального, поверніть `0`.
- Ви можете додавати лише цілі галони.
- Не включайте десяткові числа у повернене число.

# --hints--

`fuelToAdd(0, 1)` має повертати `1`.

```js
assert.equal(fuelToAdd(0, 1), 1);
```

`fuelToAdd(5, 40)` має повертати `6`.

```js
assert.equal(fuelToAdd(5, 40), 6);
```

`fuelToAdd(10, 30)` має повертати `0`.

```js
assert.equal(fuelToAdd(10, 30), 0);
```

`fuelToAdd(896, 20500)` має повертати `4520`.

```js
assert.equal(fuelToAdd(896, 20500), 4520);
```

`fuelToAdd(1000, 50000)` має повертати `12209`.

```js
assert.equal(fuelToAdd(1000, 50000), 12209);
```

# --seed--

## --seed-contents--

```js
function fuelToAdd(currentGallons, requiredLiters) {

  return currentGallons;
}
```

# --solutions--

```js
function fuelToAdd(currentGallons, requiredLiters) {
  const litersPerGallon = 3.78541;
  const currentLiters = currentGallons * litersPerGallon;
  if (currentLiters >= requiredLiters) return 0;
  const litersNeeded = requiredLiters - currentLiters;
  return Math.ceil(litersNeeded / litersPerGallon);
}
```
