---
id: 69272dcf1c24b44fd79137c3
title: "Завдання 140: SCREAMING_SNAKE_CASE"
challengeType: 28
dashedName: challenge-140
---

# --description--

Дано рядок, який представляє ім’я змінної. Поверніть ім’я змінної, конвертоване у SCREAMING_SNAKE_CASE.

Дані імена змінних будуть записані в одному з таких форматів:

- `camelCase`
- `PascalCase`
- `snake_case`
- `kebab-case`

У наведених форматах слова розділені підкресленням (`_`), дефісом (`-`) або нове слово починається з великої літери.

Щоб конвертувати у SCREAMING_SNAKE_CASE:

- Зробіть усі літери у верхньому регістрі
- Розділіть слова підкресленням (`_`)

# --hints--

`toScreamingSnakeCase("userEmail")` має повернути `"USER_EMAIL"`.

```js
assert.equal(toScreamingSnakeCase("userEmail"), "USER_EMAIL");
```

`toScreamingSnakeCase("UserPassword")` має повернути `"USER_PASSWORD"`.

```js
assert.equal(toScreamingSnakeCase("UserPassword"), "USER_PASSWORD");
```

`toScreamingSnakeCase("user_id")` має повернути `"USER_ID"`.

```js
assert.equal(toScreamingSnakeCase("user_id"), "USER_ID");
```

`toScreamingSnakeCase("user-address")` має повернути `"USER_ADDRESS"`.

```js
assert.equal(toScreamingSnakeCase("user-address"), "USER_ADDRESS");
```

`toScreamingSnakeCase("username")` має повернути `"USERNAME"`.

```js
assert.equal(toScreamingSnakeCase("username"), "USERNAME");
```

`toScreamingSnakeCase("my_variable_name")` має повернути `"MY_VARIABLE_NAME"`.

```js
assert.equal(toScreamingSnakeCase("my_variable_name"), "MY_VARIABLE_NAME");
```

# --seed--

## --seed-contents--

```js
function toScreamingSnakeCase(variableName) {

  return variableName;
}
```

# --solutions--

```js
function toScreamingSnakeCase(variableName) {
  let temp = variableName.replace(/[-_]+/g, ' ');
  temp = temp.replace(/([a-z0-9])([A-Z])/g, '$1 $2');
  const words = temp.trim().split(/\s+/);
  return words.join('_').toUpperCase();
}
```
