---
id: 69272dcf1c24b44fd79137c6
title: "Завдання 143: парсер курсиву в markdown"
challengeType: 28
dashedName: challenge-143
---

# --description--

Дано рядок, який може містити текст, виділений курсивом у Markdown. Потрібно повернути еквівалентний HTML-рядок.

- Текст курсивом у Markdown — це будь-який текст, який починається і закінчується одним символом зірочки (`*`) або одним символом підкреслення (`_`).
- Між текстом і зірочкою або підкресленням не може бути пробілів, але всередині тексту пробіли можуть бути.
- Перетворіть усі випадки курсиву на HTML-теги `i` і поверніть рядок.

Наприклад, для `"*This is italic*"` поверніть `"<i>This is italic</i>"`.

Примітка: консоль може не відображати HTML-теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseItalics("*This is italic*")` має повернути `"<i>This is italic</i>"`.

```js
assert.equal(parseItalics("*This is italic*"), "<i>This is italic</i>");
```

`parseItalics("_This is also italic_")` має повернути `"<i>This is also italic</i>"`.

```js
assert.equal(parseItalics("_This is also italic_"), "<i>This is also italic</i>");
```

`parseItalics("*This is not italic *")` має повернути `"*This is not italic *"`.

```js
assert.equal(parseItalics("*This is not italic *"), "*This is not italic *");
```

`parseItalics("_ This is also not italic_")` має повернути `"_ This is also not italic_"`.

```js
assert.equal(parseItalics("_ This is also not italic_"), "_ This is also not italic_");
```

`parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog.")` має повернути `"The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog."`.

```js
assert.equal(parseItalics("The *quick* brown fox _jumps_ over the *lazy* dog."), "The <i>quick</i> brown fox <i>jumps</i> over the <i>lazy</i> dog.");
```

# --seed--

## --seed-contents--

```js
function parseItalics(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseItalics(markdown) {
  return markdown.replace(/(\*|_)([^\s][^]*?[^\s])\1/g, '<i>$2</i>');
}
```
