---
id: 69306364df283fcaff2e1ad9
title: "Завдання 148: тиск у шинах"
challengeType: 28
dashedName: challenge-148
---

# --description--

Дано масив із чотирьох чисел, які представляють тиск у шинах вашого автомобіля в psi, та інший масив із двох чисел, які представляють мінімальний і максимальний тиск для ваших шин у барах. Поверніть масив із чотирьох рядків, що описують стан кожної шини.

- 1 бар дорівнює 14.5038 psi.

Поверніть масив із такими значеннями для кожної шини:

- `"Low"`, якщо тиск у шині нижчий за мінімально дозволений.
- `"Good"`, якщо він знаходиться між мінімальним і максимальним дозволеним.
- `"High"`, якщо він вищий за максимально дозволений.

# --hints--

`tireStatus([32, 28, 35, 29], [2, 3])` має повертати `["Good", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 29], [2, 3]), ["Good", "Low", "Good", "Low"]);
```

`tireStatus([32, 28, 35, 30], [2, 2.3])` має повертати `["Good", "Low", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([32, 28, 35, 30], [2, 2.3]), ["Good", "Low", "High", "Good"]);
```

`tireStatus([29, 26, 31, 28], [2.1, 2.5])` має повертати `["Low", "Low", "Good", "Low"]`.

```js
assert.deepEqual(tireStatus([29, 26, 31, 28], [2.1, 2.5]), ["Low", "Low", "Good", "Low"]);
```

`tireStatus([31, 31, 30, 29], [1.5, 2])` має повертати `["High", "High", "High", "Good"]`.

```js
assert.deepEqual(tireStatus([31, 31, 30, 29], [1.5, 2]), ["High", "High", "High", "Good"]);
```

`tireStatus([30, 28, 30, 29], [1.9, 2.1])` має повертати `["Good", "Good", "Good", "Good"]`.

```js
assert.deepEqual(tireStatus([30, 28, 30, 29], [1.9, 2.1]), ["Good", "Good", "Good", "Good"]);
```

# --seed--

## --seed-contents--

```js
function tireStatus(pressuresPSI, rangeBar) {

  return pressuresPSI;
}
```

# --solutions--

```js
function tireStatus(pressuresPSI, rangeBar) {
  const psiToBar = 1 / 14.5038;
  const [minBar, maxBar] = rangeBar;

  return pressuresPSI.map(psi => {
    const bar = psi * psiToBar;
    if (bar < minBar) return "Low";
    if (bar > maxBar) return "High";
    return "Good";
  });
}
```
