---
id: 69306364df283fcaff2e1ad5
title: "Завдання 144: серія успішних днів"
challengeType: 28
dashedName: challenge-144
---

# --description--

Дано масив масивів, де кожен підмасив представляє день ваших резолюцій і містить три цілі числа: кількість кроків, пройдених того дня, хвилини часу перед екраном того дня та кількість прочитаних сторінок того дня; визначте, чи ви дотримуєтеся своїх резолюцій.

- Перший підмасив — це день 1, другий — день 2 і так далі.

День вважається успішним, якщо виконуються всі три наступні цілі:

- Ви пройшли щонайменше 10 000 кроків.
- Час перед екраном не перевищував 120 хвилин.
- Ви прочитали щонайменше п’ять сторінок.

Якщо всі задані дні успішні, поверніть `"Resolution on track: N day streak."`, де `N` — кількість успішних днів.

Якщо один або більше днів не є успішними, поверніть `"Resolution failed on day X: N day streak."`, де `X` — номер першого неуспішного дня, а `N` — кількість успішних днів до першого неуспішного дня.

# --hints--

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]])` має повертати `"Resolution on track: 3 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9]]), "Resolution on track: 3 day streak.");
```

`resolutionStreak([[10000, 120, 5], [10950, 121, 11]])` має повертати `"Resolution failed on day 2: 1 day streak."`

```js
assert.equal(resolutionStreak([[10000, 120, 5], [10950, 121, 11]]), "Resolution failed on day 2: 1 day streak.");
```

`resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]])` має повертати `"Resolution failed on day 3: 2 day streak."`

```js
assert.equal(resolutionStreak([[15000, 110, 8], [12300, 60, 13], [10100, 120, 4], [9000, 125, 4]]), "Resolution failed on day 3: 2 day streak.");
```

`resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]])` має повертати `"Resolution failed on day 4: 3 day streak."`

```js
assert.equal(resolutionStreak([[11600, 76, 13], [12556, 64, 26], [10404, 32, 59], [9999, 44, 124], [7508, 23, 167], [10900, 80, 0]]), "Resolution failed on day 4: 3 day streak.");
```

`resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]])` має повертати `"Resolution on track: 10 day streak."`

```js
assert.equal(resolutionStreak([[10500, 75, 15], [11000, 90, 10], [10650, 100, 9], [10200, 60, 10], [10678, 87, 9], [12431, 67, 13], [10444, 107, 19], [10111, 95, 5], [10000, 120, 7], [11980, 101, 8]]), "Resolution on track: 10 day streak.");
```

# --seed--

## --seed-contents--

```js
function resolutionStreak(days) {

  return days;
}
```

# --solutions--

```js
function resolutionStreak(days) {
  let streak = 0;

  for (let i = 0; i < days.length; i++) {
    const [steps, screen, pages] = days[i];
    if (steps >= 10000 && screen <= 120 && pages >= 5) {
      streak++;
    } else {
      return `Resolution failed on day ${i + 1}: ${streak} day streak.`;
    }
  }

  return `Resolution on track: ${streak} day streak.`;
}
```
