---
id: 69306364df283fcaff2e1ada
title: "Завдання 149: vOwElcAsE"
challengeType: 28
dashedName: challenge-149
---

# --description--

Дано рядок. Поверніть новий рядок, у якому всі голосні перетворені на верхній регістр, а всі інші алфавітні символи — на нижній регістр.

- Голосні — це `"a"`, `"e"`, `"i"`, `"o"` і `"u"` у будь-якому регістрі.
- Неалфавітні символи повинні залишатися без змін.

# --hints--

`vowelCase("vowelcase")` має повернути `"vOwElcAsE"`.

```js
assert.equal(vowelCase("vowelcase"), "vOwElcAsE");
```

`vowelCase("coding is fun")` має повернути `"cOdIng Is fUn"`.

```js
assert.equal(vowelCase("coding is fun"), "cOdIng Is fUn");
```

`vowelCase("HELLO, world!")` має повернути `"hEllO, wOrld!"`.

```js
assert.equal(vowelCase("HELLO, world!"), "hEllO, wOrld!");
```

`vowelCase("git cherry-pick")` має повернути `"gIt chErry-pIck"`.

```js
assert.equal(vowelCase("git cherry-pick"), "gIt chErry-pIck");
```

`vowelCase("HEAD~1")` має повернути `"hEAd~1"`.

```js
assert.equal(vowelCase("HEAD~1"), "hEAd~1");
```

# --seed--

## --seed-contents--

```js
function vowelCase(str) {

  return str;
}
```

# --solutions--

```js
function vowelCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (vowels.includes(char)) {
      result += char.toUpperCase();
    } else if (/[a-zA-Z]/.test(char)) {
      result += char.toLowerCase();
    } else {
      result += char;
    }
  }

  return result;
}
```
