---
id: 69306364df283fcaff2e1ad6
title: "Завдання 145: n-те число Фібоначчі"
challengeType: 28
dashedName: challenge-145
---

# --description--

Дано ціле число `n`. Поверніть `n`-те число у послідовності Фібоначчі.

Послідовність Фібоначчі — це ряд чисел, де кожне число є сумою двох попередніх. Перші 10 чисел у послідовності: `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

# --hints--

`nthFibonacci(4)` має повернути `2`.

```js
assert.equal(nthFibonacci(4), 2);
```

`nthFibonacci(10)` має повернути `34`.

```js
assert.equal(nthFibonacci(10), 34);
```

`nthFibonacci(15)` має повернути `377`.

```js
assert.equal(nthFibonacci(15), 377);
```

`nthFibonacci(40)` має повернути `63245986`.

```js
assert.equal(nthFibonacci(40), 63245986);
```

`nthFibonacci(75)` має повернути `1304969544928657`.

```js
assert.equal(nthFibonacci(75), 1304969544928657);
```

# --seed--

## --seed-contents--

```js
function nthFibonacci(n) {

  return n;
}
```

# --solutions--

```js
function nthFibonacci(n) {
  if (n === 1) return 0;
  if (n === 2) return 1;

  let a = 0;
  let b = 1;

  for (let i = 3; i <= n; i++) {
    const next = a + b;
    a = b;
    b = next;
  }

  return b;
}
```
