---
id: 69306364df283fcaff2e1ad7
title: "Завдання 146: місце для лівші за столом"
challengeType: 28
dashedName: challenge-146
---

# --description--

Дано матрицю 4x2 (масив масивів), що представляє розсадження за столом під час їжі. Визначте, скільки місць може зайняти лівша.

- Лівша не може сидіти там, де правша сидить на місці безпосередньо ліворуч від нього.

У заданій матриці:

- `"R"` — це місце, зайняте правшею.
- `"L"` — це місце, зайняте лівшею.
- `"U"` — це вільне місце.
- Займати можна лише вільні місця.
- Місця у верхньому ряду повернені "вниз", а місця в нижньому ряду — "вверх" (як за столом), тому ліво і право відносні до орієнтації місця.
- Кутові місця мають лише одне сусіднє місце.

Наприклад, у заданій матриці:

```json
[
  ["U", "R", "U", "L"],
  ["U", "R", "R", "R"]
]
```

Місце у верхньому лівому куті зайняти не можна, бо зліва сидить правша. Інші два вільні місця можна зайняти, бо зліва немає правші.

# --hints--

`findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]])` має повернути `2`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "L"], ["U", "R", "R", "R"]]), 2);
```

`findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]])` має повернути `8`.

```js
assert.equal(findLeftHandedSeats([["U", "U", "U", "U"], ["U", "U", "U", "U"]]), 8);
```

`findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]])` має повернути `0`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "R"], ["L", "R", "R", "U"]]), 0);
```

`findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]])` має повернути `1`.

```js
assert.equal(findLeftHandedSeats([["L", "U", "R", "R"], ["L", "U", "R", "R"]]), 1);
```

`findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]])` має повернути `5`.

```js
assert.equal(findLeftHandedSeats([["U", "R", "U", "U"], ["U", "U", "L", "U"]]), 5);
```

# --seed--

## --seed-contents--

```js
function findLeftHandedSeats(table) {

  return table;
}
```

# --solutions--

```js
function findLeftHandedSeats(table) {
  let availableSeats = 0;

  const [topRow, bottomRow] = table;

  for (let i=0; i<topRow.length; i++) {
    if (topRow[i] === 'U') {
      if (i === 3) {
        availableSeats++;
      } else if (topRow[i+1] !== 'R') {
        availableSeats++;
      }
    }
  }

  for (let i=0; i<bottomRow.length; i++) {
    if (bottomRow[i] === 'U') {
      if (i === 0) {
        availableSeats++;
      } else if (bottomRow[i-1] !== 'R') {
        availableSeats++;
      }
    }
  }

  return availableSeats;
}
```
