---
id: 69306364df283fcaff2e1ad8
title: "Завдання 147: калькулятор високосного року"
challengeType: 28
dashedName: challenge-147
---

# --description--

Дано ціле число року, визначте, чи є він високосним.

Рік є високосним, якщо він задовольняє такі правила:

- Рік ділиться на 4 без остачі, і
- Рік не ділиться на 100 без остачі, якщо тільки
- Рік ділиться на 400 без остачі.

# --hints--

`isLeapYear(2024)` має повертати `true`.

```js
assert.isTrue(isLeapYear(2024));
```

`isLeapYear(2023)` має повертати `false`.

```js
assert.isFalse(isLeapYear(2023));
```

`isLeapYear(2100)` має повертати `false`.

```js
assert.isFalse(isLeapYear(2100));
```

`isLeapYear(2000)` має повертати `true`.

```js
assert.isTrue(isLeapYear(2000));
```

`isLeapYear(1999)` має повертати `false`.

```js
assert.isFalse(isLeapYear(1999));
```

`isLeapYear(2040)` має повертати `true`.

```js
assert.isTrue(isLeapYear(2040));
```

`isLeapYear(2026)` має повертати `false`.

```js
assert.isFalse(isLeapYear(2026));
```

# --seed--

## --seed-contents--

```js
function isLeapYear(year) {

  return year;
}
```

# --solutions--

```js
function isLeapYear(year) {
    return ((year % 4 === 0 && year % 100 !== 0) || (year % 400 === 0))
}
```
