---
id: 69373793f5a867f769cde13a
title: "Завдання 155: посадити врожай"
challengeType: 28
dashedName: challenge-155
---

# --description--

Дано ціле число, що представляє розмір вашого сільськогосподарського поля, та `"acres"` або `"hectares"`, які позначають одиницю виміру розміру вашого поля, а також тип культури. Визначте, скільки рослин цього типу ви можете розмістити на своєму полі.

- 1 акр дорівнює 4046,86 квадратних метрів.
- 1 гектар дорівнює 10 000 квадратних метрів.

Ось список культур, які будуть задані у вхідних даних, і скільки простору займає одна рослина:

|Культура|Простір на одну рослину|
|-|-|
|`"corn"`|1 квадратний метр|
|`"wheat"`|0,1 квадратного метра|
|`"soybeans"`|0,5 квадратного метра|
|`"tomatoes"`|0,25 квадратного метра|
|`"lettuce"`|0,2 квадратного метра|

Поверніть кількість рослин, які помістяться на полі, округлену вниз до найближчого цілого числа рослин.

# --hints--

`getNumberOfPlants(1, "acres", "corn")` має повертати `4046`.

```js
assert.equal(getNumberOfPlants(1, "acres", "corn"), 4046);
```

`getNumberOfPlants(2, "hectares", "lettuce")` має повертати `100000`.

```js
assert.equal(getNumberOfPlants(2, "hectares", "lettuce"), 100000);
```

`getNumberOfPlants(20, "acres", "soybeans")` має повертати `161874`.

```js
assert.equal(getNumberOfPlants(20, "acres", "soybeans"), 161874);
```

`getNumberOfPlants(3.75, "hectares", "tomatoes")` має повертати `150000`.

```js
assert.equal(getNumberOfPlants(3.75, "hectares", "tomatoes"), 150000);
```

`getNumberOfPlants(16.75, "acres", "tomatoes")` має повертати `271139`.

```js
assert.equal(getNumberOfPlants(16.75, "acres", "tomatoes"), 271139);
```

# --seed--

## --seed-contents--

```js
function getNumberOfPlants(fieldSize, unit, crop) {

  return fieldSize;
}
```

# --solutions--

```js
function getNumberOfPlants(fieldSize, unit, crop) {
  const cropSpace = {
    corn: 1,
    wheat: 0.1,
    soybeans: 0.5,
    tomatoes: 0.25,
    lettuce: 0.2
  };

  let fieldInSqM = unit === "acres" ? fieldSize * 4046.86 : fieldSize * 10000;

  return Math.floor(fieldInSqM / cropSpace[crop.toLowerCase()]);
}
```
