---
id: 69373793f5a867f769cde137
title: "Завдання 152: кругле просте число"
challengeType: 28
dashedName: challenge-152
---

# --description--

Дано ціле число, визначте, чи є воно круглим простим числом.

Кругле просте число — це ціле число, у якого всі оберти цифр самі є простими числами.

Наприклад, `197` є круглим простим числом, тому що всі оберти його цифр: `197`, `971` і `719` — прості числа.

# --hints--

`isCircularPrime(197)` має повернути `true`.

```js
assert.isTrue(isCircularPrime(197));
```

`isCircularPrime(23)` має повернути `false`.

```js
assert.isFalse(isCircularPrime(23));
```

`isCircularPrime(13)` має повернути `true`.

```js
assert.isTrue(isCircularPrime(13));
```

`isCircularPrime(89)` має повернути `false`.

```js
assert.isFalse(isCircularPrime(89));
```

`isCircularPrime(1193)` має повернути `true`.

```js
assert.isTrue(isCircularPrime(1193));
```

# --seed--

## --seed-contents--

```js
function isCircularPrime(n) {

  return n;
}
```

# --solutions--

```js
function isPrime(n) {
  if (n < 2) return false;
  for (let i = 2, sqrt = Math.floor(Math.sqrt(n)); i <= sqrt; i++) {
    if (n % i === 0) return false;
  }
  return true;
}

function rotations(num) {
  const str = num.toString();
  const result = [];
  for (let i = 0; i < str.length; i++) {
    result.push(str.slice(i) + str.slice(0, i));
  }
  return result.map(Number);
}

function isCircularPrime(num) {
  const rots = rotations(num);
  return rots.every(isPrime);
}
```
