---
id: 69373793f5a867f769cde136
title: "Завдання 151: відсортований масив?"
challengeType: 28
dashedName: challenge-151
---

# --description--

Дано масив чисел. Визначте, чи відсортовані числа за зростанням, за спаданням або ні за якимось із цих порядків.

Якщо заданий масив:

- Відсортований за зростанням (від найменшого до найбільшого), поверніть `"Ascending"`.
- Відсортований за спаданням (від найбільшого до найменшого), поверніть `"Descending"`.
- Не відсортований ні за зростанням, ні за спаданням, поверніть `"Not sorted"`.

# --hints--

`isSorted([1, 2, 3, 4, 5])` має повертати `"Ascending"`.

```js
assert.equal(isSorted([1, 2, 3, 4, 5]), "Ascending");
```

`isSorted([10, 8, 6, 4, 2])` має повертати `"Descending"`.

```js
assert.equal(isSorted([10, 8, 6, 4, 2]), "Descending");
```

`isSorted([1, 3, 2, 4, 5])` має повертати `"Not sorted"`.

```js
assert.equal(isSorted([1, 3, 2, 4, 5]), "Not sorted");
```

`isSorted([3.14, 2.71, 1.61, 0.57])` має повертати `"Descending"`.

```js
assert.equal(isSorted([3.14, 2.71, 1.61, 0.57]), "Descending");
```

`isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9])` має повертати `"Ascending"`.

```js
assert.equal(isSorted([12.3, 23.4, 34.5, 45.6, 56.7, 67.8, 78.9]), "Ascending");
```

`isSorted([0.4, 0.5, 0.3])` має повертати `"Not sorted"`.

```js
assert.equal(isSorted([0.4, 0.5, 0.3]), "Not sorted");
```

# --seed--

## --seed-contents--

```js
function isSorted(arr) {

  return arr;
}
```

# --solutions--

```js
function isSorted(arr) {
  let ascending = true;
  let descending = true;

  for (let i = 1; i < arr.length; i++) {
    if (arr[i] < arr[i - 1]) ascending = false;
    if (arr[i] > arr[i - 1]) descending = false;
  }

  if (ascending) return "Ascending";
  if (descending) return "Descending";
  return "Not sorted";
}
```
