---
id: 69373793f5a867f769cde135
title: "Завдання 150: парсер markdown невпорядкованого списку"
challengeType: 28
dashedName: challenge-150
---

# --description--

Дано рядок з дійсним markdown невпорядкованим списком, поверніть еквівалентний HTML рядок.

Невпорядкований список складається з одного або кількох елементів списку. Дійсний елемент списку розміщується на окремій лінії і:

- Починається з дефіса (`"-"`), за яким іде
- Принаймні один пробіл, а потім
- Текст елемента списку.

Список подається як один рядок з новими лініями, розділеними символом нового рядка (`"\n"`). Не включайте символи нового рядка у текст елемента.

Обгорніть кожен елемент списку в HTML теги `li`, а весь список елементів — у теги `ul`.

Наприклад, для `"- Item A\n- Item B"` поверніть `"<ul><li>Item A</li><li>Item B</li></ul>"`.

Примітка: консоль може не відображати HTML теги в рядках при виводі повідомлень. Перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseUnorderedList("- Item A\n- Item B")` має повернути `"<ul><li>Item A</li><li>Item B</li></ul>"`.

```js
assert.equal(parseUnorderedList("- Item A\n- Item B"), "<ul><li>Item A</li><li>Item B</li></ul>");
```

`parseUnorderedList("-  JavaScript\n-  Python")` має повернути `"<ul><li>JavaScript</li><li>Python</li></ul>"`.

```js
assert.equal(parseUnorderedList("-  JavaScript\n-  Python"), "<ul><li>JavaScript</li><li>Python</li></ul>");
```

`parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla")` має повернути `"<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>"`.

```js
assert.equal(parseUnorderedList("- 2 C Flour\n- 1/2 C Sugar\n- 1 Tsp Vanilla"), "<ul><li>2 C Flour</li><li>1/2 C Sugar</li><li>1 Tsp Vanilla</li></ul>");
```

`parseUnorderedList("- A-1\n- A-2\n- B-1")` має повернути `"<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>"`.

```js
assert.equal(parseUnorderedList("- A-1\n- A-2\n- B-1"), "<ul><li>A-1</li><li>A-2</li><li>B-1</li></ul>");
```

# --seed--

## --seed-contents--

```js
function parseUnorderedList(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseUnorderedList(markdown) {
  const lines = markdown.split("\n");
  const listItems = [];
  
  for (let line of lines) {
    listItems.push(`<li>${line.slice(1).trim()}</li>`);
  }
  
  return `<ul>${listItems.join("")}</ul>`;
}
```
