---
id: 69373793f5a867f769cde139
title: "Завдання 154: Пар для лунки"
challengeType: 28
dashedName: challenge-154
---

# --description--

Дано два цілі числа: пар для гольф-лунки та кількість штрихів, які гравець зробив на цій лунці. Поверніть результат гравця, використовуючи терміни гольфу.

Поверніть:

- `"Hole in one!"`, якщо знадобився один штрих.
- `"Eagle"`, якщо знадобилось на два штрихи менше за пар.
- `"Birdie"`, якщо знадобився на один штрих менше за пар.
- `"Par"`, якщо кількість штрихів дорівнює пар.
- `"Bogey"`, якщо знадобився на один штрих більше за пар.
- `"Double bogey"`, якщо знадобилось на два штрихи більше за пар.

# --hints--

`golfScore(3, 3)` має повертати `"Par"`.

```js
assert.equal(golfScore(3, 3), "Par");
```

`golfScore(4, 3)` має повертати `"Birdie"`.

```js
assert.equal(golfScore(4, 3), "Birdie");
```

`golfScore(3, 1)` має повертати `"Hole in one!"`.

```js
assert.equal(golfScore(3, 1), "Hole in one!");
```

`golfScore(5, 7)` має повертати `"Double bogey"`.

```js
assert.equal(golfScore(5, 7), "Double bogey");
```

`golfScore(4, 5)` має повертати `"Bogey"`.

```js
assert.equal(golfScore(4, 5), "Bogey");
```

`golfScore(5, 3)` має повертати `"Eagle"`.

```js
assert.equal(golfScore(5, 3), "Eagle");
```

# --seed--

## --seed-contents--

```js
function golfScore(par, strokes) {

  return par;
}
```

# --solutions--

```js
function golfScore(par, strokes) {
  if (strokes === 1) return "Hole in one!";

  const diff = strokes - par;

  if (diff === -2) return "Eagle";
  if (diff === -1) return "Birdie";
  if (diff === 0) return "Par";
  if (diff === 1) return "Bogey";
  if (diff === 2) return "Double bogey";
}
```
