---
id: 6939b873185d8e00d453563c
title: "Завдання 157: парсер посилань у markdown"
challengeType: 28
dashedName: challenge-157
---

# --description--

Дано рядок посилання у форматі Markdown, поверніть еквівалентний рядок у форматі HTML.

Зображення у Markdown має такий формат: `"[link_text](link_url)"`. Поверніть рядок HTML-тегу `a` з множиною `href`, встановленою у `link_url`, а `link_text` — як вміст тегу.

Наприклад, для `"[freeCodeCamp](https://freecodecamp.org/)"` поверніть `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`;

Примітка: консоль може не відображати HTML-теги у рядках при виводі повідомлень — перевірте консоль браузера, щоб побачити логи з тегами.

# --hints--

`parseLink("[freeCodeCamp](https://freecodecamp.org/)")` має повертати `'<a href="https://freecodecamp.org/">freeCodeCamp</a>'`.

```js
assert.equal(parseLink("[freeCodeCamp](https://freecodecamp.org/)"), '<a href="https://freecodecamp.org/">freeCodeCamp</a>');
```

`parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)")` має повертати `'<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>'`.

```js
assert.equal(parseLink("[Donate to our charity.](https://www.freecodecamp.org/donate/)"), '<a href="https://www.freecodecamp.org/donate/">Donate to our charity.</a>');
```

`parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)")` має повертати `'<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>'`.

```js
assert.equal(parseLink("[Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.](https://github.com/freeCodeCamp/freeCodeCamp/)"), '<a href="https://github.com/freeCodeCamp/freeCodeCamp/">Contribute to our repository at https://github.com/freeCodeCamp/freeCodeCamp.</a>');
```

# --seed--

## --seed-contents--

```js
function parseLink(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseLink(markdown) {
  const match = markdown.match(/\[(.*?)\]\((.*?)\)/);

  const [, text, url] = match;
  return `<a href="${url}">${text}</a>`;
}
```
