---
id: 6939b873185d8e00d453563d
title: "Завдання 158: обмін значень у масиві"
challengeType: 28
dashedName: challenge-158
---

# --description--

Дано масив із двома значеннями, поверніть масив із обміняними значеннями.

Наприклад, для `["A", "B"]` поверніть `["B", "A"]`.

# --hints--

`arraySwap(["A", "B"])` має повернути `["B", "A"]`.

```js
assert.deepEqual(arraySwap(["A", "B"]), ["B", "A"]);
```

`arraySwap([25, 20])` має повернути `[20, 25]`.

```js
assert.deepEqual(arraySwap([25, 20]), [20, 25]);
```

`arraySwap([true, false])` має повернути `[false, true]`.

```js
assert.deepEqual(arraySwap([true, false]), [false, true]);
```

`arraySwap(["1", 1])` має повернути `[1, "1"]`.

```js
assert.deepEqual(arraySwap(["1", 1]), [1, "1"]);
```

# --seed--

## --seed-contents--

```js
function arraySwap(arr) {

  return arr;
}
```

# --solutions--

```js
function arraySwap(arr) {

  return [arr[1], arr[0]];
}
```
