---
id: 6939b873185d8e00d453563f
title: "Завдання 160: ходи конем"
challengeType: 28
dashedName: challenge-160
---

# --description--

За заданою позицією коня на шахівниці поверніть число дійсних клітин, на які конь може ходити.

Стандартна шахівниця має розмір 8x8, з колонками, позначеними `A` до `H` (зліва направо), і рядами, позначеними `1` до `8` (знизу вгору). Вона виглядає так:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Конь ходить у формі літери "L": дві клітини в одному напрямку (горизонтально або вертикально) і одна клітина в перпендикулярному напрямку.

Це означає, що конь може ходити до восьми можливих позицій, але менше, якщо він знаходиться біля країв дошки. Наприклад, якщо конь стоїть на `A1`, він може ходити лише на `B3` або `C2`.

# --hints--

`knightMoves("A1")` має повернути `2`.

```js
assert.equal(knightMoves("A1"), 2);
```

`knightMoves("D4")` має повернути `8`.

```js
assert.equal(knightMoves("D4"), 8);
```

`knightMoves("G6")` має повернути `6`.

```js
assert.equal(knightMoves("G6"), 6);
```

`knightMoves("B8")` має повернути `3`.

```js
assert.equal(knightMoves("B8"), 3);
```

`knightMoves("H3")` має повернути `4`.

```js
assert.equal(knightMoves("H3"), 4);
```

# --seed--

## --seed-contents--

```js
function knightMoves(position) {

  return position;
}
```

# --solutions--

```js
function knightMoves(position) {
  const col = position[0].charCodeAt(0) - 65;
  const row = parseInt(position[1]) - 1;

  const moves = [
    [2, 1], [2, -1],
    [-2, 1], [-2, -1],
    [1, 2], [1, -2],
    [-1, 2], [-1, -2]
  ];

  let valid = 0;

  for (const [dx, dy] of moves) {
    const newCol = col + dx;
    const newRow = row + dy;

    if (newCol >= 0 && newCol < 8 && newRow >= 0 && newRow < 8) {
      valid++;
    }
  }

  return valid;
}
```
