---
id: 6939b873185d8e00d453563b
title: "Завдання 156: непарне чи парне?"
challengeType: 28
dashedName: challenge-156
---

# --description--

Дано додатнє ціле число. Поверніть `"Odd"`, якщо це непарне число, і `"Even"`, якщо парне.

# --hints--

`oddOrEven(1)` має повертати `"Odd"`.

```js
assert.equal(oddOrEven(1), "Odd");
```

`oddOrEven(2)` має повертати `"Even"`.

```js
assert.equal(oddOrEven(2), "Even");
```

`oddOrEven(13)` має повертати `"Odd"`.

```js
assert.equal(oddOrEven(13), "Odd");
```

`oddOrEven(196)` має повертати `"Even"`.

```js
assert.equal(oddOrEven(196), "Even");
```

`oddOrEven(123456789)` має повертати `"Odd"`.

```js
assert.equal(oddOrEven(123456789), "Odd");
```

# --seed--

## --seed-contents--

```js
function oddOrEven(n) {

  return n;
}
```

# --solutions--

```js
function oddOrEven(n) {
  return n % 2 === 0 ? "Even" : "Odd";
}
```
