---
id: 6939b873185d8e00d453563e
title: "Завдання 159: цілий гіпотенуз"
challengeType: 28
dashedName: challenge-159
---

# --description--

Дано два натуральні числа, які представляють довжини двох катетів (двох коротких сторін) прямокутного трикутника. Визначте, чи є гіпотенуза цілим числом.

Довжина гіпотенузи обчислюється шляхом додавання квадратів довжин двох катетів, а потім взяття квадратного кореня від цієї суми (a<sup>2</sup> + b<sup>2</sup> = c<sup>2</sup>).

# --hints--

`isIntegerHypotenuse(3, 4)` має повертати `true`.

```js
assert.isTrue(isIntegerHypotenuse(3, 4));
```

`isIntegerHypotenuse(2, 3)` має повертати `false`.

```js
assert.isFalse(isIntegerHypotenuse(2, 3));
```

`isIntegerHypotenuse(5, 12)` має повертати `true`.

```js
assert.isTrue(isIntegerHypotenuse(5, 12));
```

`isIntegerHypotenuse(10, 10)` має повертати `false`.

```js
assert.isFalse(isIntegerHypotenuse(10, 10));
```

`isIntegerHypotenuse(780, 1040)` має повертати `true`.

```js
assert.isTrue(isIntegerHypotenuse(780, 1040));
```

`isIntegerHypotenuse(250, 333)` має повертати `false`.

```js
assert.isFalse(isIntegerHypotenuse(250, 333));
```

# --seed--

## --seed-contents--

```js
function isIntegerHypotenuse(a, b) {

  return a;
}
```

# --solutions--

```js
function isIntegerHypotenuse(a, b) {
  const sum = a * a + b * b;
  const c = Math.floor(Math.sqrt(sum));
  return c * c === sum;
}
```
