---
id: 694596b0585c11170ac7c7fc
title: "Завдання 164: парсер рядкового коду в markdown"
challengeType: 28
dashedName: challenge-164
---

# --description--

Дано рядок Markdown, який містить один або кілька рядкових блоків коду. Поверніть еквівалентний HTML-рядок.

Рядкові блоки коду в Markdown використовують одинарний зворотний апостроф (``` ` ```) на початку та в кінці тексту коду.

Поверніть заданий рядок, у якому всі блоки коду перетворені на HTML-теги `code`.

Наприклад, для рядка ```"Use `let` to declare the variable."```, return `"Use <code>let</code> to declare the variable."`.

Примітка: консоль може не відображати HTML-теги в рядках під час виведення повідомлень. Перевірте консоль браузера, щоб побачити логи з включеними тегами.

# --hints--

```parseInlineCode("Use `let` to declare the variable.")``` should return `"Use <code>let</code> to declare the variable."`.

```js
assert.equal(parseInlineCode("Use `let` to declare the variable."), "Use <code>let</code> to declare the variable.");
```

```parseInlineCode("Use `let` or `const` to declare a variable.")``` should return `"Use <code>let</code> or <code>const</code> to declare a variable."`.

```js
assert.equal(parseInlineCode("Use `let` or `const` to declare a variable."), "Use <code>let</code> or <code>const</code> to declare a variable.");
```

```parseInlineCode("Run `npm install` then `npm start`.")``` should return `"Run <code>npm install</code> then <code>npm start</code>."`.

```js
assert.equal(parseInlineCode("Run `npm install` then `npm start`."), "Run <code>npm install</code> then <code>npm start</code>.");
```

# --seed--

## --seed-contents--

```js
function parseInlineCode(markdown) {

  return markdown;
}
```

# --solutions--

```js
function parseInlineCode(markdown) {
  return markdown.replace(/`([^`]+)`/g, "<code>$1</code>");
}
```
