---
id: 694596b0585c11170ac7c7fb
title: "Завдання 163: регістр приголосних"
challengeType: 28
dashedName: challenge-163
---

# --description--

Дано рядок, який представляє ім’я змінної. Перетворіть його на регістр приголосних за такими правилами:

- Усі приголосні мають бути перетворені на верхній регістр.
- Усі голосні (`a`, `e`, `i`, `o`, `u` у будь-якому регістрі) мають бути перетворені на нижній регістр.
- Усі дефіси (`-`) мають бути замінені на підкреслення (`_`).

# --hints--

`toConsonantCase("helloworld")` має повернути `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("helloworld"), "HeLLoWoRLD");
```

`toConsonantCase("HELLOWORLD")` має повернути `"HeLLoWoRLD"`.

```js
assert.equal(toConsonantCase("HELLOWORLD"), "HeLLoWoRLD");
```

`toConsonantCase("_hElLO-WOrlD-")` має повернути `"_HeLLo_WoRLD_"`.

```js
assert.equal(toConsonantCase("_hElLO-WOrlD-"), "_HeLLo_WoRLD_");
```

`toConsonantCase("_~-generic_~-variable_~-name_~-here-~_")` має повернути `"_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_"`.

```js
assert.equal(toConsonantCase("_~-generic_~-variable_~-name_~-here-~_"), "_~_GeNeRiC_~_VaRiaBLe_~_NaMe_~_HeRe_~_");
```

# --seed--

## --seed-contents--

```js
function toConsonantCase(str) {

  return str;
}
```

# --solutions--

```js
function toConsonantCase(str) {
  const vowels = "aeiouAEIOU";
  let result = "";

  for (let char of str) {
    if (char === "-") {
      result += "_";
    } else if (/[a-zA-Z]/.test(char)) {
      if (vowels.includes(char)) {
        result += char.toLowerCase();
      } else {
        result += char.toUpperCase();
      }
    } else {
      result += char;
    }
  }

  return result;
}
```
