---
id: 694596b0585c11170ac7c7fa
title: "Завдання 162: споживання енергії"
challengeType: 28
dashedName: challenge-162
---

# --description--

Враховуючи число калорій, спалених під час тренування, та число ват-годин, використаних вашими електронними пристроями під час цього тренування, визначте, що використало більше енергії.

Щоб порівняти їх, перетворіть обидва значення в джоулі, використовуючи такі конверсії:

- 1 калорія дорівнює 4184 джоулів.
- 1 ват-година дорівнює 3600 джоулів.

Поверніть:

- `"Workout"`, якщо тренування використало більше енергії.
- `"Devices"`, якщо пристрій використав більше енергії.
- `"Equal"`, якщо обидва використали однакову кількість енергії.

# --hints--

`compareEnergy(250, 50)` має повертати `"Workout"`.

```js
assert.equal(compareEnergy(250, 50), "Workout");
```

`compareEnergy(100, 200)` має повертати `"Devices"`.

```js
assert.equal(compareEnergy(100, 200), "Devices");
```

`compareEnergy(450, 523)` має повертати `"Equal"`.

```js
assert.equal(compareEnergy(450, 523), "Equal");
```

`compareEnergy(300, 75)` має повертати `"Workout"`.

```js
assert.equal(compareEnergy(300, 75), "Workout");
```

`compareEnergy(200, 250)` має повертати `"Devices"`.

```js
assert.equal(compareEnergy(200, 250), "Devices");
```

`compareEnergy(900, 1046)` має повертати `"Equal"`.

```js
assert.equal(compareEnergy(900, 1046), "Equal");
```

# --seed--

## --seed-contents--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {

  return caloriesBurned;
}
```

# --solutions--

```js
function compareEnergy(caloriesBurned, wattHoursUsed) {
  const workoutEnergy = caloriesBurned * 4184;
  const deviceEnergy = wattHoursUsed * 3600;

  console.log(workoutEnergy);
  console.log(deviceEnergy);
  if (workoutEnergy > deviceEnergy) return "Workout";
  if (deviceEnergy > workoutEnergy) return "Devices";
  return "Equal";
}
```
