---
id: 696655d24b614176d4c9b78d
title: "Завдання 170: непарний чи парний день"
challengeType: 28
dashedName: challenge-170
---

# --description--

Дано часову позначку (число мілісекунд від початку епохи Unix), поверніть:

- `"odd"`, якщо число дня місяця для цієї часової позначки непарне.
- `"even"`, якщо число дня місяця для цієї часової позначки парне.

Наприклад, для `1769472000000`, часової позначки 27 січня 2026 року, поверніть `"odd"`, бо число дня (`27`) є непарним числом.

Примітка: Часова позначка задана в мілісекундах, і слід використовувати дату в часовому поясі UTC, а не у вашому локальному часі.

# --hints--

`oddOrEvenDay(1769472000000)` має повертати `"odd"`.

```js
assert.equal(oddOrEvenDay(1769472000000), "odd");
```

`oddOrEvenDay(1769444440000)` має повертати `"even"`.

```js
assert.equal(oddOrEvenDay(1769444440000), "even");
```

`oddOrEvenDay(6739456780000)` має повертати `"odd"`.

```js
assert.equal(oddOrEvenDay(6739456780000), "odd");
```

`oddOrEvenDay(1)` має повертати `"odd"`.

```js
assert.equal(oddOrEvenDay(1), "odd");
```

`oddOrEvenDay(86400000)` має повертати `"even"`.

```js
assert.equal(oddOrEvenDay(86400000), "even");
```

# --seed--

## --seed-contents--

```js
function oddOrEvenDay(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function oddOrEvenDay(timestamp) {
  const date = new Date(timestamp);
  const day = date.getUTCDate();

  return day % 2 === 0 ? "even" : "odd";
}
```
