---
id: 696655d24b614176d4c9b78b
title: "Завдання 168: масштабоване зображення"
challengeType: 28
dashedName: challenge-168
---

# --description--

Дано рядок, який представляє ширину та висоту зображення, і число для масштабування зображення. Поверніть масштабовані ширину та висоту.

- Вхідний рядок має формат `"WxH"`. Наприклад, `"800x600"`.
- Масштаб — це число, на яке потрібно помножити ширину та висоту.

Поверніть масштабовані розміри у тому самому форматі `"WxH"`.

# --hints--

`scaleImage("800x600", 2)` має повернути `"1600x1200"`.

```js
assert.equal(scaleImage("800x600", 2), "1600x1200");
```

`scaleImage("100x100", 10)` має повернути `"1000x1000"`.

```js
assert.equal(scaleImage("100x100", 10), "1000x1000");
```

`scaleImage("1024x768", 0.5)` має повернути `"512x384"`.

```js
assert.equal(scaleImage("1024x768", 0.5), "512x384");
```

`scaleImage("300x200", 1.5)` має повернути `"450x300"`.

```js
assert.equal(scaleImage("300x200", 1.5), "450x300");
```

# --seed--

## --seed-contents--

```js
function scaleImage(size, scale) {

  return size;
}
```

# --solutions--

```js
function scaleImage(size, scale) {
  const [width, height] = size.split("x").map(Number);

  const newWidth = width * scale;
  const newHeight = height * scale;

  return `${newWidth}x${newHeight}`;
}
```
