---
id: 696655d24b614176d4c9b78c
title: "Завдання 169: міні fizzbuzz"
challengeType: 28
dashedName: challenge-169
---

# --description--

Дано ціле число, поверніть рядок відповідно до таких правил:

- Якщо число ділиться на 3, поверніть `"Fizz"`.
- Якщо число ділиться на 5, поверніть `"Buzz"`.
- Якщо число ділиться і на 3, і на 5, поверніть `"FizzBuzz"`.
- Інакше поверніть задане число у вигляді рядка.

# --hints--

`fizzBuzzMini(3)` має повертати `"Fizz"`.

```js
assert.equal(fizzBuzzMini(3), "Fizz");
```

`fizzBuzzMini(4)` має повертати `"4"`.

```js
assert.strictEqual(fizzBuzzMini(4), "4");
```

`fizzBuzzMini(35)` має повертати `"Buzz"`.

```js
assert.equal(fizzBuzzMini(35), "Buzz");
```

`fizzBuzzMini(75)` має повертати `"FizzBuzz"`.

```js
assert.equal(fizzBuzzMini(75), "FizzBuzz");
```

`fizzBuzzMini(98)` має повертати `"98"`.

```js
assert.strictEqual(fizzBuzzMini(98), "98");
```

# --seed--

## --seed-contents--

```js
function fizzBuzzMini(n) {

  return n;
}
```

# --solutions--

```js
function fizzBuzzMini(n) {
  if (n % 3 === 0 && n % 5 === 0) return "FizzBuzz";
  if (n % 3 === 0) return "Fizz";
  if (n % 5 === 0) return "Buzz";
  return String(n);
}
```
