---
id: 696655d24b614176d4c9b789
title: "Завдання 166: перевірка шістнадцяткового коду"
challengeType: 28
dashedName: challenge-166
---

# --description--

Дано рядок. Визначте, чи є він дійсним CSS шістнадцятковим кольором. Дійсний CSS шістнадцятковий колір повинен:

- Починатися з `#`, і
- містити після нього або 3, або 6 шістнадцяткових символів.

Шістнадцяткові символи — це числа `0` до `9` та літери `a` до `f` (регістр не має значення).

# --hints--

`isValidHex("#123")` має повертати `true`.

```js
assert.isTrue(isValidHex("#123"));
```

`isValidHex("#123abc")` має повертати `true`.

```js
assert.isTrue(isValidHex("#123abc"));
```

`isValidHex("#ABCDEF")` має повертати `true`.

```js
assert.isTrue(isValidHex("#ABCDEF"));
```

`isValidHex("#0a1B2c")` має повертати `true`.

```js
assert.isTrue(isValidHex("#0a1B2c"));
```

`isValidHex("#12G")` має повертати `false`.

```js
assert.isFalse(isValidHex("#12G"));
```

`isValidHex("#1234567")` має повертати `false`.

```js
assert.isFalse(isValidHex("#1234567"));
```

`isValidHex("#12 3")` має повертати `false`.

```js
assert.isFalse(isValidHex("#12 3"));
```

`isValidHex("fff")` має повертати `false`.

```js
assert.isFalse(isValidHex("fff"));
```

# --seed--

## --seed-contents--

```js
function isValidHex(str) {

  return str;
}
```

# --solutions--

```js
function isValidHex(str) {
  if (str[0] !== "#") return false;

  const hex = str.slice(1);
  if (hex.length !== 3 && hex.length !== 6) return false;

  for (let char of hex) {
    if (!/[0-9a-fA-F]/.test(char)) {
      return false;
    }
  }

  return true;
}
```
