---
id: 69738771fb5a7b8b24cca29f
title: "Завдання 173: допустимі ходи пішака"
challengeType: 28
dashedName: challenge-173
---

# --description--

Враховуючи позицію одного з ваших пішаків на шахівниці, поверніть масив усіх допустимих клітинок, на які він може ходити, у зростаючому порядку.

Стандартна шахівниця має розмір 8x8, з колонками, позначеними `A` до `H` (зліва направо), і рядами, позначеними `1` до `8` (знизу вгору). Вона виглядає так:

|**A8**|**B8**|**C8**|**D8**|**E8**|**F8**|**G8**|**H8**|
|-|-|-|-|-|-|-|-|
|**A7**|**B7**|**C7**|**D7**|**E7**|**F7**|**G7**|**H7**|
|**A6**|**B6**|**C6**|**D6**|**E6**|**F6**|**G6**|**H6**|
|**A5**|**B5**|**C5**|**D5**|**E5**|**F5**|**G5**|**H5**|
|**A4**|**B4**|**C4**|**D4**|**E4**|**F4**|**G4**|**H4**|
|**A3**|**B3**|**C3**|**D3**|**E3**|**F3**|**G3**|**H3**|
|**A2**|**B2**|**C2**|**D2**|**E2**|**F2**|**G2**|**H2**|
|**A1**|**B1**|**C1**|**D1**|**E1**|**F1**|**G1**|**H1**|

Для цього завдання:

- Ви гравець, що стоїть внизу дошки.
- Пішаки можуть рухатися лише на одну клітинку "вгору".
- Якщо пішак знаходиться на стартовому ряді (ряд 2), він може рухатися на одну або дві клітинки вгору.

Наприклад, якщо задано `"D4"`, поверніть `["D5"]` — єдину клітинку, на яку може ходити ваш пішак.
Якщо задано `"B2"`, поверніть `["B3", "B4"]`, бо він на стартовому ряді, і результат має бути у зростаючому порядку.

# --hints--

`findPawnMoves("D4")` має повернути `["D5"]`.

```js
assert.deepEqual(findPawnMoves("D4"), ["D5"]);
```

`findPawnMoves("B2")` має повернути `["B3", "B4"]`.

```js
assert.deepEqual(findPawnMoves("B2"), ["B3", "B4"]);
```

`findPawnMoves("A7")` має повернути `["A8"]`.

```js
assert.deepEqual(findPawnMoves("A7"), ["A8"]);
```

`findPawnMoves("G2")` має повернути `["G3", "G4"]`.

```js
assert.deepEqual(findPawnMoves("G2"), ["G3", "G4"]);
```

`findPawnMoves("E3")` має повернути `["E4"]`.

```js
assert.deepEqual(findPawnMoves("E3"), ["E4"]);
```

# --seed--

## --seed-contents--

```js
function findPawnMoves(position) {

  return position;
}

```

# --solutions--

```js
function findPawnMoves(position) {
  const column = position[0].toUpperCase();
  const row = parseInt(position[1]);
  const moves = [];

  if (row >= 8) return moves;

  moves.push(`${column}${row + 1}`);

  if (row === 2) moves.push(`${column}${row + 2}`);

  return moves;
}
```
