---
id: 69738771fb5a7b8b24cca2a2
title: "Завдання 176: День бабака"
challengeType: 28
dashedName: challenge-176
---

# --description--

Сьогодні День бабака, коли бабак передбачає погоду, спостерігаючи, чи бачить він свою тінь.

Враховуючи значення, що представляє появу бабака, поверніть правильне передбачення:

- Якщо задане значення — булевий `true` (бабак побачив свою тінь), поверніть `"Looks like we'll have six more weeks of winter."`.
- Якщо значення — булевий `false` (бабак не побачив свою тінь), поверніть `"It's going to be an early spring."`.
- Якщо значення будь-яке інше (бабак не з’явився), поверніть `"No prediction this year."`.

# --hints--

`groundhogDayPrediction(true)` має повертати `"Looks like we'll have six more weeks of winter."`.

```js
assert.equal(groundhogDayPrediction(true), "Looks like we'll have six more weeks of winter.");
```

`groundhogDayPrediction(false)` має повертати `"It's going to be an early spring."`.

```js
assert.equal(groundhogDayPrediction(false), "It's going to be an early spring.");
```

`groundhogDayPrediction(null)` має повертати `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(null), "No prediction this year.");
```

`groundhogDayPrediction(" ")` має повертати `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction(" "), "No prediction this year.");
```

`groundhogDayPrediction("true")` має повертати `"No prediction this year."`.

```js
assert.equal(groundhogDayPrediction("true"), "No prediction this year.");
```

# --seed--

## --seed-contents--

```js
function groundhogDayPrediction(appearance) {

  return appearance;
}
```

# --solutions--

```js
function groundhogDayPrediction(appearance) {
  if (appearance === true) return "Looks like we'll have six more weeks of winter.";
  if (appearance === false) return "It's going to be an early spring.";
  return "No prediction this year.";
}
```
