---
id: 69738771fb5a7b8b24cca2a3
title: "Завдання 177: дзеркальний рядок"
challengeType: 28
dashedName: challenge-177
---

# --description--

Дано рядок. Поверніть новий рядок, який складається з даного рядка з доданою в кінці перевернутою копією самого себе.

# --hints--

`mirror("freeCodeCamp")` має повернути `"freeCodeCamppmaCedoCeerf"`.

```js
assert.equal(mirror("freeCodeCamp"), "freeCodeCamppmaCedoCeerf");
```

`mirror("RaceCar")` має повернути `"RaceCarraCecaR"`.

```js
assert.equal(mirror("RaceCar"), "RaceCarraCecaR");
```

`mirror("helloworld")` має повернути `"helloworlddlrowolleh"`.

```js
assert.equal(mirror("helloworld"), "helloworlddlrowolleh");
```

`mirror("The quick brown fox...")` має повернути `"The quick brown fox......xof nworb kciuq ehT"`.

```js
assert.equal(mirror("The quick brown fox..."), "The quick brown fox......xof nworb kciuq ehT");
```

# --seed--

## --seed-contents--

```js
function mirror(str) {

  return str;
}
```

# --solutions--

```js
function mirror(str) {
  const reversed = str.split("").reverse().join("");
  return str + reversed;
}
```
