---
id: 69738771fb5a7b8b24cca2a4
title: "Завдання 178: скоротити текст"
challengeType: 28
dashedName: challenge-178
---

# --description--

Дано рядок. Поверніть його без змін, якщо він містить 20 символів або менше. Якщо він довший за 20 символів, скоротіть його до перших 17 символів і додайте `"..."` у кінець (щоб загальна довжина була 20 символів), а потім поверніть результат.

# --hints--

`truncateText("Hello, world!")` має повертати `"Hello, world!"`.

```js
assert.equal(truncateText("Hello, world!"), "Hello, world!");
```

`truncateText("This string should get truncated.")` має повертати `"This string shoul..."`.

```js
assert.equal(truncateText("This string should get truncated."), "This string shoul...");
```

`truncateText("Exactly twenty chars")` має повертати `"Exactly twenty chars"`.

```js
assert.equal(truncateText("Exactly twenty chars"), "Exactly twenty chars");
```

`truncateText(".....................")` має повертати `"...................."`.

```js
assert.equal(truncateText("....................."), "....................");
```

# --seed--

## --seed-contents--

```js
function truncateText(text) {

  return text;
}
```

# --solutions--

```js
function truncateText(text) {
  if (text.length <= 20) return text;
  return text.slice(0, 17) + "...";
}
```
