---
id: 69738771fb5a7b8b24cca2a0
title: "Завдання 174: пошук знаку зодіаку"
challengeType: 28
dashedName: challenge-174
---

# --description--

Дано рядок дати у форматі `"YYYY-MM-DD"`. Поверніть знак зодіаку для цієї дати, використовуючи таку діаграму:

| Діапазон дат | Знак зодіаку |
| - | - |
| 21 березня - 19 квітня | `"Aries"` |
| 20 квітня - 20 травня | `"Taurus"` |
| 21 травня - 20 червня | `"Gemini"` |
| 21 червня - 22 липня | `"Cancer"` |
| 23 липня - 22 серпня | `"Leo"` |
| 23 серпня - 22 вересня | `"Virgo"` |
| 23 вересня - 22 жовтня | `"Libra"` |
| 23 жовтня - 21 листопада | `"Scorpio"` |
| 22 листопада - 21 грудня | `"Sagittarius"` |
| 22 грудня - 19 січня | `"Capricorn"` |
| 20 січня - 18 лютого | `"Aquarius"` |
| 19 лютого - 20 березня | `"Pisces"` |

- Знаки зодіаку визначаються лише за місяцем і днем, рік можна ігнорувати.

# --hints--

`getSign("2026-01-31")` має повернути `"Aquarius"`.

```js
assert.equal(getSign("2026-01-31"), "Aquarius");
```

`getSign("2001-06-10")` має повернути `"Gemini"`.

```js
assert.equal(getSign("2001-06-10"), "Gemini");
```

`getSign("1985-09-07")` має повернути `"Virgo"`.

```js
assert.equal(getSign("1985-09-07"), "Virgo");
```

`getSign("2023-03-19")` має повернути `"Pisces"`.

```js
assert.equal(getSign("2023-03-19"), "Pisces");
```

`getSign("2045-11-05")` має повернути `"Scorpio"`.

```js
assert.equal(getSign("2045-11-05"), "Scorpio");
```

`getSign("1985-12-06")` має повернути `"Sagittarius"`.

```js
assert.equal(getSign("1985-12-06"), "Sagittarius");
```

`getSign("2025-12-30")` має повернути `"Capricorn"`.

```js
assert.equal(getSign("2025-12-30"), "Capricorn");
```

`getSign("2018-10-08")` має повернути `"Libra"`.

```js
assert.equal(getSign("2018-10-08"), "Libra");
```

`getSign("1958-05-04")` має повернути `"Taurus"`.

```js
assert.equal(getSign("1958-05-04"), "Taurus");
```

# --seed--

## --seed-contents--

```js
function getSign(dateStr) {

  return dateStr;
}
```

# --solutions--

```js
function getSign(dateStr) {
  const [, monthStr, dayStr] = dateStr.split("-");
  const month = Number(monthStr);
  const day = Number(dayStr);

  if ((month === 3 && day >= 21) || (month === 4 && day <= 19)) return "Aries";
  if ((month === 4 && day >= 20) || (month === 5 && day <= 20)) return "Taurus";
  if ((month === 5 && day >= 21) || (month === 6 && day <= 20)) return "Gemini";
  if ((month === 6 && day >= 21) || (month === 7 && day <= 22)) return "Cancer";
  if ((month === 7 && day >= 23) || (month === 8 && day <= 22)) return "Leo";
  if ((month === 8 && day >= 23) || (month === 9 && day <= 22)) return "Virgo";
  if ((month === 9 && day >= 23) || (month === 10 && day <= 22)) return "Libra";
  if ((month === 10 && day >= 23) || (month === 11 && day <= 21)) return "Scorpio";
  if ((month === 11 && day >= 22) || (month === 12 && day <= 21)) return "Sagittarius";
  if ((month === 12 && day >= 22) || (month === 1 && day <= 19)) return "Capricorn";
  if ((month === 1 && day >= 20) || (month === 2 && day <= 18)) return "Aquarius";
  if ((month === 2 && day >= 19) || (month === 3 && day <= 20)) return "Pisces";
}
```
