---
id: 697a49e9860d24853adef680
title: "Завдання 195: зимові ігри 2026, день 16: керлінг"
challengeType: 28
dashedName: challenge-195
---

# --description--

Дано матрицю 5x5, що представляє "дім" наприкінці раунду керлінгу. Визначте, яка команда набирає очки і скільки саме.

Макет:

- Центральна клітинка (індекс `[2, 2]`) — це "кнопка".
- 8 клітинок безпосередньо навколо кнопки представляють кільце 1.
- 16 клітинок на зовнішньому краю дому представляють кільце 2.

У заданій матриці:

- `"."` означає порожній пробіл.
- `"R"` означає пробіл із червоним каменем.
- `"Y"` означає пробіл із жовтим каменем.

Правила нарахування очок:

- За раунд очки може набрати лише одна команда.
- Команда, камінь якої найближчий до кнопки, отримує очки.
- Команда, що набирає очки, отримує 1 бал за кожен свій камінь, який ближчий до кнопки, ніж найближчий камінь суперника.
- Чим менший номер кільця, тим ближче камінь до центру.
- Якщо найближчі камені обох команд знаходяться на однаковій відстані від центру, жодна команда не набирає очок.

Поверніть:

- рядок у форматі `"team: number_of_points"`, наприклад: `"R: 2"`.
- або `"No points awarded"`, якщо жодна команда не набрала очок.

Наприклад, для:

```js
[
  [".", ".", "R", ".", "."],
  [".", "R", ".", ".", "."],
  ["Y", ".", ".", ".", "."],
  [".", "R", ".", ".", "."],
  [".", ".", ".", ".", "."]
]
```

Поверніть `"R: 2"`. Два червоні камені в кільці 1 мають однакову найближчу відстань і є єдиними каменями, що ближчі за найближчий жовтий.

# --hints--

`scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]])` має повернути `"R: 2"`.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", "R", ".", ".", "."], ["Y", ".", ".", ".", "."], [".", "R", ".", ".", "."], [".", ".", ".", ".", "."]]), "R: 2");
```

`scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]])` має повернути `"Y: 3"`.

```js
assert.equal(scoreCurling([[".", ".", "R", ".", "."], [".", ".", ".", ".", "."], [".", ".", "Y", ".", "R"], [".", ".", "Y", "Y", "."], [".", "Y", "R", "R", "."]]), "Y: 3");
```

`scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]])` має повернути `"No points awarded"`.

```js
assert.equal(scoreCurling([[".", "R", "Y", ".", "."], ["Y", ".", ".", ".", "."], [".", ".", ".", ".", "."], [".", "Y", "R", "Y", "."], [".", ".", "R", "R", "."]]), "No points awarded");
```

`scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]])` має повернути `"R: 4"`.

```js
assert.equal(scoreCurling([[".", "Y", "Y", ".", "."], ["Y", ".", ".", "R", "."], [".", ".", "R", ".", "."], [".", ".", "R", "R", "."], [".", "Y", "R", "Y", "."]]), "R: 4");
```

`scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]])` має повернути `"Y: 1"`.

```js
assert.equal(scoreCurling([["Y", "Y", "Y", "Y", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "R", "Y", "R", "Y"], ["Y", "R", "R", "R", "Y"], ["Y", "Y", "Y", "Y", "Y"]]), "Y: 1");
```

`scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]])` має повернути `"No points awarded"`.

```js
assert.equal(scoreCurling([["Y", "R", "Y", "R", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", ".", "Y"], ["R", ".", ".", ".", "R"], ["Y", ".", ".", "R", "Y"]]), "No points awarded");
```

# --seed--

## --seed-contents--

```js
function scoreCurling(house) {

  return house;
}
```

# --solutions--

```js
function scoreCurling(house) {
  const stones = { R: [], Y: [] };

  for (let i = 0; i < 5; i++) {
    for (let j = 0; j < 5; j++) {
      const cell = house[i][j];
      if (cell === 'R' || cell === 'Y') {
        const distance = Math.max(Math.abs(i - 2), Math.abs(j - 2));
        stones[cell].push(distance);
      }
    }
  }

  const redClosest = stones.R.length > 0 ? Math.min(...stones.R) : Infinity;
  const yellowClosest = stones.Y.length > 0 ? Math.min(...stones.Y) : Infinity;

  if (redClosest === Infinity && yellowClosest === Infinity) {
    return "No points awarded";
  }

  if (redClosest === yellowClosest) {
    return "No points awarded";
  }

let winner, opponentClosest, winningStones;
  if (redClosest < yellowClosest) {
    winner = 'R';
    opponentClosest = yellowClosest;
    winningStones = stones.R;
  } else {
    winner = 'Y';
    opponentClosest = redClosest;
    winningStones = stones.Y;
  }

  const points = winningStones.filter(d => d < opponentClosest).length;

  return points === 0 ? "No points awarded" : `${winner}: ${points}`;
}
```
