---
id: 697a49e9860d24853adef67c
title: "Завдання 191: зимові ігри 2026, день 12: бобслей"
challengeType: 28
dashedName: challenge-191
---

# --description--

Дано масив, який представляє вагу спортсменів у команді з бобслею, та число, що представляє вагу бобслею. Визначте, чи команда має право брати участь у змаганні.

- Довжина масиву визначає розмір команди: 1, 2 або 4 особи.
- Усі наведені значення ваги вказані в кілограмах (кг).

Бобслей (сам сани) повинен мати мінімальну вагу:

- 162 кг для команди з 1 особи
- 170 кг для команди з 2 осіб
- 210 кг для команди з 4 осіб

Загальна вага бобслею (спортсмени плюс сани) не повинна перевищувати:

- 247 кг для команди з 1 особи
- 390 кг для команди з 2 осіб
- 630 кг для команди з 4 осіб

Поверніть `"Eligible"`, якщо команда відповідає всім вимогам, або `"Not Eligible"`, якщо команда не відповідає одній або кільком вимогам.

# --hints--

`checkEligibility([78], 165)` має повертати `"Eligible"`.

```js
assert.equal(checkEligibility([78], 165), "Eligible");
```

`checkEligibility([80], 160)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 160), "Not Eligible");
```

`checkEligibility([80], 170)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([80], 170), "Not Eligible");
```

`checkEligibility([85, 90], 170)` має повертати `"Eligible"`.

```js
assert.equal(checkEligibility([85, 90], 170), "Eligible");
```

`checkEligibility([85, 95], 168)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([85, 95], 168), "Not Eligible");
```

`checkEligibility([112, 97], 185)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([112, 97], 185), "Not Eligible");
```

`checkEligibility([110, 102, 90, 106], 222)` має повертати `"Eligible"`.

```js
assert.equal(checkEligibility([110, 102, 90, 106], 222), "Eligible");
```

`checkEligibility([106, 99, 90, 88], 205)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 90, 88], 205), "Not Eligible");
```

`checkEligibility([106, 99, 103, 96], 227)` має повертати `"Not Eligible"`.

```js
assert.equal(checkEligibility([106, 99, 103, 96], 227), "Not Eligible");
```

# --seed--

## --seed-contents--

```js
function checkEligibility(athleteWeights, sledWeight) {

  return athleteWeights;
}
```

# --solutions--

```js
function checkEligibility(athleteWeights, sledWeight) {
  const teamSize = athleteWeights.length;

  const rules = {
    1: { minSled: 162, maxTotal: 247 },
    2: { minSled: 170, maxTotal: 390 },
    4: { minSled: 210, maxTotal: 630 }
  };

  const { minSled, maxTotal } = rules[teamSize];
  const athleteTotal = athleteWeights.reduce((sum, w) => sum + w, 0);
  const totalWeight = athleteTotal + sledWeight;

  if (sledWeight < minSled) {
    return "Not Eligible";
  }

  if (totalWeight > maxTotal) {
    return "Not Eligible";
  }

  return "Eligible";
}
```
