---
id: 697a49e6ff50d756c9b6935d
title: "Завдання 180: Зимові ігри, день 1: день відкриття"
challengeType: 28
dashedName: challenge-180
---

# --description--

Наступні 17 днів принесуть вам завдання з програмування, натхненні Зимовими іграми.

Для першого завдання вам дають дволітерний код країни, і потрібно повернути емодзі прапорця цієї країни.

Використовуйте цей список:

| Країна                    | Код            | Прапорець |
| -                         | -:             | :--:      |
| Albania                   |            `"AL"` | `"🇦🇱"` |
| Andorra                   |            `"AD"` | `"🇦🇩"` |
| Argentina                 |            `"AR"` | `"🇦🇷"` |
| Armenia                   |            `"AM"` | `"🇦🇲"` |
| Australia                 |            `"AU"` | `"🇦🇺"` |
| Austria                   |            `"AT"` | `"🇦🇹"` |
| Azerbaijan                |            `"AZ"` | `"🇦🇿"` |
| Belgium                   |            `"BE"` | `"🇧🇪"` |
| Benin                     |            `"BJ"` | `"🇧🇯"` |
| Bolivia                   |            `"BO"` | `"🇧🇴"` |
| Bosnia and Herzegovina    |            `"BA"` | `"🇧🇦"` |
| Brazil                    |            `"BR"` | `"🇧🇷"` |
| Bulgaria                  |            `"BG"` | `"🇧🇬"` |
| Canada                    |            `"CA"` | `"🇨🇦"` |
| Chile                     |            `"CL"` | `"🇨🇱"` |
| China                     |            `"CN"` | `"🇨🇳"` |
| Colombia                  |            `"CO"` | `"🇨🇴"` |
| Croatia                   |            `"HR"` | `"🇭🇷"` |
| Cyprus                    |            `"CY"` | `"🇨🇾"` |
| Czech Republic            |            `"CZ"` | `"🇨🇿"` |
| Denmark                   |            `"DK"` | `"🇩🇰"` |
| Ecuador                   |            `"EC"` | `"🇪🇨"` |
| Eritrea                   |            `"ER"` | `"🇪🇷"` |
| Estonia                   |            `"EE"` | `"🇪🇪"` |
| Finland                   |            `"FI"` | `"🇫🇮"` |
| France                    |            `"FR"` | `"🇫🇷"` |
| Georgia                   |            `"GE"` | `"🇬🇪"` |
| Germany                   |            `"DE"` | `"🇩🇪"` |
| Great Britain             |            `"GB"` | `"🇬🇧"` |
| Greece                    |            `"GR"` | `"🇬🇷"` |
| Guinea-Bissau             |            `"GW"` | `"🇬🇼"` |
| Haiti                     |            `"HT"` | `"🇭🇹"` |
| Hong Kong                 |            `"HK"` | `"🇭🇰"` |
| Hungary                   |            `"HU"` | `"🇭🇺"` |
| Iceland                   |            `"IS"` | `"🇮🇸"` |
| India                     |            `"IN"` | `"🇮🇳"` |
| Iran                      |            `"IR"` | `"🇮🇷"` |
| Ireland                   |            `"IE"` | `"🇮🇪"` |
| Israel                    |            `"IL"` | `"🇮🇱"` |
| Italy                     |            `"IT"` | `"🇮🇹"` |
| Jamaica                   |            `"JM"` | `"🇯🇲"` |
| Japan                     |            `"JP"` | `"🇯🇵"` |
| Kazakhstan                |            `"KZ"` | `"🇰🇿"` |
| Kenya                     |            `"KE"` | `"🇰🇪"` |
| Kosovo                    |            `"XK"` | `"🇽🇰"` |
| Kyrgyzstan                |            `"KG"` | `"🇰🇬"` |
| Latvia                    |            `"LV"` | `"🇱🇻"` |
| Lebanon                   |            `"LB"` | `"🇱🇧"` |
| Liechtenstein             |            `"LI"` | `"🇱🇮"` |
| Lithuania                 |            `"LT"` | `"🇱🇹"` |
| Luxembourg                |            `"LU"` | `"🇱🇺"` |
| Madagascar                |            `"MG"` | `"🇲🇬"` |
| Malaysia                  |            `"MY"` | `"🇲🇾"` |
| Malta                     |            `"MT"` | `"🇲🇹"` |
| Mexico                    |            `"MX"` | `"🇲🇽"` |
| Moldova                   |            `"MD"` | `"🇲🇩"` |
| Monaco                    |            `"MC"` | `"🇲🇨"` |
| Mongolia                  |            `"MN"` | `"🇲🇳"` |
| Montenegro                |            `"ME"` | `"🇲🇪"` |
| Morocco                   |            `"MA"` | `"🇲🇦"` |
| Netherlands               |            `"NL"` | `"🇳🇱"` |
| New Zealand               |            `"NZ"` | `"🇳🇿"` |
| Nigeria                   |            `"NG"` | `"🇳🇬"` |
| North Macedonia           |            `"MK"` | `"🇲🇰"` |
| Norway                    |            `"NO"` | `"🇳🇴"` |
| Pakistan                  |            `"PK"` | `"🇵🇰"` |
| Philippines               |            `"PH"` | `"🇵🇭"` |
| Poland                    |            `"PL"` | `"🇵🇱"` |
| Portugal                  |            `"PT"` | `"🇵🇹"` |
| Puerto Rico               |            `"PR"` | `"🇵🇷"` |
| Romania                   |            `"RO"` | `"🇷🇴"` |
| San Marino                |            `"SM"` | `"🇸🇲"` |
| Saudi Arabia              |            `"SA"` | `"🇸🇦"` |
| Serbia                    |            `"RS"` | `"🇷🇸"` |
| Singapore                 |            `"SG"` | `"🇸🇬"` |
| Slovakia                  |            `"SK"` | `"🇸🇰"` |
| Slovenia                  |            `"SI"` | `"🇸🇮"` |
| South Africa              |            `"ZA"` | `"🇿🇦"` |
| South Korea               |            `"KR"` | `"🇰🇷"` |
| Spain                     |            `"ES"` | `"🇪🇸"` |
| Sweden                    |            `"SE"` | `"🇸🇪"` |
| Switzerland               |            `"CH"` | `"🇨🇭"` |
| Thailand                  |            `"TH"` | `"🇹🇭"` |
| Trinidad & Tobago         |            `"TT"` | `"🇹🇹"` |
| Turkey                    |            `"TR"` | `"🇹🇷"` |
| Ukraine                   |            `"UA"` | `"🇺🇦"` |
| United Arab Emirates      |            `"AE"` | `"🇦🇪"` |
| United States             |            `"US"` | `"🇺🇸"` |
| Uruguay                   |            `"UY"` | `"🇺🇾"` |
| Uzbekistan                |            `"UZ"` | `"🇺🇿"` |
| Venezuela                 |            `"VE"` | `"🇻🇪"` |

# --hints--

`getFlag("AL")` має повертати `"🇦🇱"`.

```js
assert.equal(getFlag("AL"), "🇦🇱");
```

`getFlag("AD")` має повертати `"🇦🇩"`.

```js
assert.equal(getFlag("AD"), "🇦🇩");
```

`getFlag("AR")` має повертати `"🇦🇷"`.

```js
assert.equal(getFlag("AR"), "🇦🇷")
```

`getFlag("AM")` має повертати `"🇦🇲"`.

```js
assert.equal(getFlag("AM"), "🇦🇲")
```

`getFlag("AU")` має повертати `"🇦🇺"`.

```js
assert.equal(getFlag("AU"), "🇦🇺")
```

`getFlag("AT")` має повертати `"🇦🇹"`.

```js
assert.equal(getFlag("AT"), "🇦🇹")
```

`getFlag("AZ")` має повертати `"🇦🇿"`.

```js
assert.equal(getFlag("AZ"), "🇦🇿")
```

`getFlag("BE")` має повертати `"🇧🇪"`.

```js
assert.equal(getFlag("BE"), "🇧🇪")
```

`getFlag("BJ")` має повертати `"🇧🇯"`.

```js
assert.equal(getFlag("BJ"), "🇧🇯")
```

`getFlag("BO")` має повертати `"🇧🇴"`.

```js
assert.equal(getFlag("BO"), "🇧🇴")
```

`getFlag("BA")` має повертати `"🇧🇦"`.

```js
assert.equal(getFlag("BA"), "🇧🇦")
```

`getFlag("BR")` має повертати `"🇧🇷"`.

```js
assert.equal(getFlag("BR"), "🇧🇷")
```

`getFlag("BG")` має повертати `"🇧🇬"`.

```js
assert.equal(getFlag("BG"), "🇧🇬")
```

`getFlag("CA")` має повертати `"🇨🇦"`.

```js
assert.equal(getFlag("CA"), "🇨🇦")
```

`getFlag("CL")` має повертати `"🇨🇱"`.

```js
assert.equal(getFlag("CL"), "🇨🇱")
```

`getFlag("CN")` має повертати `"🇨🇳"`.

```js
assert.equal(getFlag("CN"), "🇨🇳")
```

`getFlag("CO")` має повертати `"🇨🇴"`.

```js
assert.equal(getFlag("CO"), "🇨🇴")
```

`getFlag("HR")` має повертати `"🇭🇷"`.

```js
assert.equal(getFlag("HR"), "🇭🇷")
```

`getFlag("CY")` має повертати `"🇨🇾"`.

```js
assert.equal(getFlag("CY"), "🇨🇾")
```

`getFlag("CZ")` має повертати `"🇨🇿"`.

```js
assert.equal(getFlag("CZ"), "🇨🇿")
```

`getFlag("DK")` має повертати `"🇩🇰"`.

```js
assert.equal(getFlag("DK"), "🇩🇰")
```

`getFlag("EC")` має повертати `"🇪🇨"`.

```js
assert.equal(getFlag("EC"), "🇪🇨")
```

`getFlag("ER")` має повертати `"🇪🇷"`.

```js
assert.equal(getFlag("ER"), "🇪🇷")
```

`getFlag("EE")` має повертати `"🇪🇪"`.

```js
assert.equal(getFlag("EE"), "🇪🇪")
```

`getFlag("FI")` має повертати `"🇫🇮"`.

```js
assert.equal(getFlag("FI"), "🇫🇮")
```

`getFlag("FR")` має повертати `"🇫🇷"`.

```js
assert.equal(getFlag("FR"), "🇫🇷")
```

`getFlag("GE")` має повертати `"🇬🇪"`.

```js
assert.equal(getFlag("GE"), "🇬🇪")
```

`getFlag("DE")` має повертати `"🇩🇪"`.

```js
assert.equal(getFlag("DE"), "🇩🇪")
```

`getFlag("GB")` має повертати `"🇬🇧"`.

```js
assert.equal(getFlag("GB"), "🇬🇧")
```

`getFlag("GR")` має повертати `"🇬🇷"`.

```js
assert.equal(getFlag("GR"), "🇬🇷")
```

`getFlag("GW")` має повертати `"🇬🇼"`.

```js
assert.equal(getFlag("GW"), "🇬🇼")
```

`getFlag("HT")` має повертати `"🇭🇹"`.

```js
assert.equal(getFlag("HT"), "🇭🇹")
```

`getFlag("HK")` має повертати `"🇭🇰"`.

```js
assert.equal(getFlag("HK"), "🇭🇰")
```

`getFlag("HU")` має повертати `"🇭🇺"`.

```js
assert.equal(getFlag("HU"), "🇭🇺")
```

`getFlag("IS")` має повертати `"🇮🇸"`.

```js
assert.equal(getFlag("IS"), "🇮🇸")
```

`getFlag("IN")` має повертати `"🇮🇳"`.

```js
assert.equal(getFlag("IN"), "🇮🇳")
```

`getFlag("IR")` має повертати `"🇮🇷"`.

```js
assert.equal(getFlag("IR"), "🇮🇷")
```

`getFlag("IE")` має повертати `"🇮🇪"`.

```js
assert.equal(getFlag("IE"), "🇮🇪")
```

`getFlag("IL")` має повертати `"🇮🇱"`.

```js
assert.equal(getFlag("IL"), "🇮🇱")
```

`getFlag("IT")` має повертати `"🇮🇹"`.

```js
assert.equal(getFlag("IT"), "🇮🇹")
```

`getFlag("JM")` має повертати `"🇯🇲"`.

```js
assert.equal(getFlag("JM"), "🇯🇲")
```

`getFlag("JP")` має повертати `"🇯🇵"`.

```js
assert.equal(getFlag("JP"), "🇯🇵")
```

`getFlag("KZ")` має повертати `"🇰🇿"`.

```js
assert.equal(getFlag("KZ"), "🇰🇿")
```

`getFlag("KE")` має повертати `"🇰🇪"`.

```js
assert.equal(getFlag("KE"), "🇰🇪")
```

`getFlag("XK")` має повертати `"🇽🇰"`.

```js
assert.equal(getFlag("XK"), "🇽🇰")
```

`getFlag("KG")` має повертати `"🇰🇬"`.

```js
assert.equal(getFlag("KG"), "🇰🇬")
```

`getFlag("LV")` має повертати `"🇱🇻"`.

```js
assert.equal(getFlag("LV"), "🇱🇻")
```

`getFlag("LB")` має повертати `"🇱🇧"`.

```js
assert.equal(getFlag("LB"), "🇱🇧")
```

`getFlag("LI")` має повертати `"🇱🇮"`.

```js
assert.equal(getFlag("LI"), "🇱🇮")
```

`getFlag("LT")` має повертати `"🇱🇹"`.

```js
assert.equal(getFlag("LT"), "🇱🇹")
```

`getFlag("LU")` має повертати `"🇱🇺"`.

```js
assert.equal(getFlag("LU"), "🇱🇺")
```

`getFlag("MG")` має повертати `"🇲🇬"`.

```js
assert.equal(getFlag("MG"), "🇲🇬")
```

`getFlag("MY")` має повертати `"🇲🇾"`.

```js
assert.equal(getFlag("MY"), "🇲🇾")
```

`getFlag("MT")` має повертати `"🇲🇹"`.

```js
assert.equal(getFlag("MT"), "🇲🇹")
```

`getFlag("MX")` має повертати `"🇲🇽"`.

```js
assert.equal(getFlag("MX"), "🇲🇽")
```

`getFlag("MD")` має повертати `"🇲🇩"`.

```js
assert.equal(getFlag("MD"), "🇲🇩")
```

`getFlag("MC")` має повертати `"🇲🇨"`.

```js
assert.equal(getFlag("MC"), "🇲🇨")
```

`getFlag("MN")` має повертати `"🇲🇳"`.

```js
assert.equal(getFlag("MN"), "🇲🇳")
```

`getFlag("ME")` має повертати `"🇲🇪"`.

```js
assert.equal(getFlag("ME"), "🇲🇪")
```

`getFlag("MA")` має повертати `"🇲🇦"`.

```js
assert.equal(getFlag("MA"), "🇲🇦")
```

`getFlag("NL")` має повертати `"🇳🇱"`.

```js
assert.equal(getFlag("NL"), "🇳🇱")
```

`getFlag("NZ")` має повертати `"🇳🇿"`.

```js
assert.equal(getFlag("NZ"), "🇳🇿")
```

`getFlag("NG")` має повертати `"🇳🇬"`.

```js
assert.equal(getFlag("NG"), "🇳🇬")
```

`getFlag("MK")` має повертати `"🇲🇰"`.

```js
assert.equal(getFlag("MK"), "🇲🇰")
```

`getFlag("NO")` має повертати `"🇳🇴"`.

```js
assert.equal(getFlag("NO"), "🇳🇴")
```

`getFlag("PK")` має повертати `"🇵🇰"`.

```js
assert.equal(getFlag("PK"), "🇵🇰")
```

`getFlag("PH")` має повертати `"🇵🇭"`.

```js
assert.equal(getFlag("PH"), "🇵🇭")
```

`getFlag("PL")` має повертати `"🇵🇱"`.

```js
assert.equal(getFlag("PL"), "🇵🇱")
```

`getFlag("PT")` має повертати `"🇵🇹"`.

```js
assert.equal(getFlag("PT"), "🇵🇹")
```

`getFlag("PR")` має повертати `"🇵🇷"`.

```js
assert.equal(getFlag("PR"), "🇵🇷")
```

`getFlag("RO")` має повертати `"🇷🇴"`.

```js
assert.equal(getFlag("RO"), "🇷🇴")
```

`getFlag("SM")` має повертати `"🇸🇲"`.

```js
assert.equal(getFlag("SM"), "🇸🇲")
```

`getFlag("SA")` має повертати `"🇸🇦"`.

```js
assert.equal(getFlag("SA"), "🇸🇦")
```

`getFlag("RS")` має повертати `"🇷🇸"`.

```js
assert.equal(getFlag("RS"), "🇷🇸")
```

`getFlag("SG")` має повертати `"🇸🇬"`.

```js
assert.equal(getFlag("SG"), "🇸🇬")
```

`getFlag("SK")` має повертати `"🇸🇰"`.

```js
assert.equal(getFlag("SK"), "🇸🇰")
```

`getFlag("SI")` має повертати `"🇸🇮"`.

```js
assert.equal(getFlag("SI"), "🇸🇮")
```

`getFlag("ZA")` має повертати `"🇿🇦"`.

```js
assert.equal(getFlag("ZA"), "🇿🇦")
```

`getFlag("KR")` має повертати `"🇰🇷"`.

```js
assert.equal(getFlag("KR"), "🇰🇷")
```

`getFlag("ES")` має повертати `"🇪🇸"`.

```js
assert.equal(getFlag("ES"), "🇪🇸")
```

`getFlag("SE")` має повертати `"🇸🇪"`.

```js
assert.equal(getFlag("SE"), "🇸🇪")
```

`getFlag("CH")` має повертати `"🇨🇭"`.

```js
assert.equal(getFlag("CH"), "🇨🇭")
```

`getFlag("TH")` має повертати `"🇹🇭"`.

```js
assert.equal(getFlag("TH"), "🇹🇭")
```

`getFlag("TT")` має повертати `"🇹🇹"`.

```js
assert.equal(getFlag("TT"), "🇹🇹")
```

`getFlag("TR")` має повертати `"🇹🇷"`.

```js
assert.equal(getFlag("TR"), "🇹🇷")
```

`getFlag("UA")` має повертати `"🇺🇦"`.

```js
assert.equal(getFlag("UA"), "🇺🇦")
```

`getFlag("AE")` має повертати `"🇦🇪"`.

```js
assert.equal(getFlag("AE"), "🇦🇪")
```

`getFlag("US")` має повертати `"🇺🇸"`.

```js
assert.equal(getFlag("US"), "🇺🇸")
```

`getFlag("UY")` має повертати `"🇺🇾"`.

```js
assert.equal(getFlag("UY"), "🇺🇾")
```

`getFlag("UZ")` має повертати `"🇺🇿"`.

```js
assert.equal(getFlag("UZ"), "🇺🇿")
```

`getFlag("VE")` має повертати `"🇻🇪"`.

```js
assert.equal(getFlag("VE"), "🇻🇪")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
