---
id: 697a49e6ff50d756c9b6935f
title: "Завдання 182: зимові ігри 2026, день 3: біатлон"
challengeType: 28
dashedName: challenge-182
---

# --description--

Дано масив цілих чисел, де кожне значення представляє кількість влучень у ціль за один раунд біатлону. Поверніть загальну штрафну дистанцію, яку спортсмен має пробігти.

- Кожен раунд складається з 5 мішеней.
- За кожну пропущену мішень спортсмен має пробігти штрафний цикл довжиною 150 метрів.

# --hints--

`calculatePenaltyDistance([4, 4])` має повернути `300`.

```js
assert.equal(calculatePenaltyDistance([4, 4]), 300);
```

`calculatePenaltyDistance([5, 5])` має повернути `0`.

```js
assert.equal(calculatePenaltyDistance([5, 5]), 0);
```

`calculatePenaltyDistance([4, 5, 3, 5])` має повернути `450`.

```js
assert.equal(calculatePenaltyDistance([4, 5, 3, 5]), 450);
```

`calculatePenaltyDistance([5, 4, 5, 5])` має повернути `150`.

```js
assert.equal(calculatePenaltyDistance([5, 4, 5, 5]), 150);
```

`calculatePenaltyDistance([4, 3, 0, 3])` має повернути `1500`.

```js
assert.equal(calculatePenaltyDistance([4, 3, 0, 3]), 1500);
```

# --seed--

## --seed-contents--

```js
function calculatePenaltyDistance(rounds) {

  return rounds;
}
```

# --solutions--

```js
function calculatePenaltyDistance(rounds) {
  let totalPenalty = 0;

  for (const hits of rounds) {
    const misses = 5 - hits;
    totalPenalty += misses * 150;
  }

  return totalPenalty;
}
```
