---
id: 697a49e6ff50d756c9b69364
title: "Завдання 187: зимові ігри 2026, день 8: санний спорт"
challengeType: 28
dashedName: challenge-187
---

# --description--

Дано масив із п’яти чисел, кожне з яких представляє час (у секундах), за який саночник пройшов сегмент траси. Визначте, у якому сегменті була найвища швидкість і яке це було значення.

Траса поділена на такі сегменти:

- Сегмент 1: 320 метрів
- Сегмент 2: 280 метрів
- Сегмент 3: 350 метрів
- Сегмент 4: 300 метрів
- Сегмент 5: 250 метрів

Перше значення в масиві відповідає часу для сегмента 1, друге — для сегмента 2 і так далі.

Щоб обчислити швидкість (у метрах за секунду) для сегмента, поділіть відстань на час.

Поверніть `"The luger's fastest speed was X m/s on segment Y."`, де `X` — найвища швидкість, округлена до двох десяткових знаків, а `Y` — номер сегмента, у якому зафіксована найвища швидкість.

# --hints--

`getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128])` має повертати `"The luger's fastest speed was 35.07 m/s on segment 5."`

```js
assert.equal(getFastestSpeed([9.523, 8.234, 10.012, 9.001, 7.128]), "The luger's fastest speed was 35.07 m/s on segment 5.");
```

`getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284])` має повертати `"The luger's fastest speed was 37.75 m/s on segment 2."`

```js
assert.equal(getFastestSpeed([9.381, 7.417, 9.912, 8.815, 7.284]), "The luger's fastest speed was 37.75 m/s on segment 2.");
```

`getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912])` має повертати `"The luger's fastest speed was 38.49 m/s on segment 3."`

```js
assert.equal(getFastestSpeed([8.890, 7.601, 9.093, 8.392, 6.912]), "The luger's fastest speed was 38.49 m/s on segment 3.");
```

`getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840])` має повертати `"The luger's fastest speed was 37.69 m/s on segment 1."`

```js
assert.equal(getFastestSpeed([8.490, 7.732, 10.103, 8.489, 6.840]), "The luger's fastest speed was 37.69 m/s on segment 1.");
```

`getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508])` має повертати `"The luger's fastest speed was 39.24 m/s on segment 4."`

```js
assert.equal(getFastestSpeed([8.204, 7.230, 9.673, 7.645, 6.508]), "The luger's fastest speed was 39.24 m/s on segment 4.");
```

# --seed--

## --seed-contents--

```js
function getFastestSpeed(times) {

  return times;
}
```

# --solutions--

```js
function getFastestSpeed(times) {
  const distances = [320, 280, 350, 300, 250];
  let fastestSpeed = 0;
  let fastestSegment = 0;

  for (let i = 0; i < times.length; i++) {
    const speed = distances[i] / times[i];
    if (speed > fastestSpeed) {
      fastestSpeed = speed;
      fastestSegment = i + 1;
    }
  }

  return `The luger's fastest speed was ${fastestSpeed.toFixed(2)} m/s on segment ${fastestSegment}.`;
}
```
