---
id: 697a49e6ff50d756c9b69366
title: "Завдання 189: зимові ігри 2026, день 10: гірськолижний спорт"
challengeType: 28
dashedName: challenge-189
---

# --description--

За заданим вертикальним перепадом, горизонтальною відстанню та типом схилу визначте рейтинг складності схилу.

Щоб визначити рейтинг:

- Обчисліть крутість схилу, поділивши перепад на відстань.
- Потім обчисліть скориговану крутість залежно від типу схилу:
  - `"Downhill"`: помножте крутість на 1.2
  - `"Slalom"`: помножте крутість на 0.9
  - `"Giant Slalom"`: помножте крутість на 1.0

Поверніть:

- `"Green"`, якщо скоригована крутість менша або дорівнює 0.1
- `"Blue"`, якщо скоригована крутість більша за 0.1 і менша або дорівнює 0.25
- `"Black"`, якщо скоригована крутість більша за 0.25

# --hints--

`getHillRating(95, 900, "Slalom")` має повертати `"Green"`.

```js
assert.equal(getHillRating(95, 900, "Slalom"), "Green");
```

`getHillRating(620, 2800, "Downhill")` має повертати `"Black"`.

```js
assert.equal(getHillRating(620, 2800, "Downhill"), "Black");
```

`getHillRating(420, 1680, "Giant Slalom")` має повертати `"Blue"`.

```js
assert.equal(getHillRating(420, 1680, "Giant Slalom"), "Blue");
```

`getHillRating(250, 3000, "Downhill")` має повертати `"Green"`.

```js
assert.equal(getHillRating(250, 3000, "Downhill"), "Green");
```

`getHillRating(110, 900, "Slalom")` має повертати `"Blue"`.

```js
assert.equal(getHillRating(110, 900, "Slalom"), "Blue");
```

`getHillRating(380, 1500, "Giant Slalom")` має повертати `"Black"`.

```js
assert.equal(getHillRating(380, 1500, "Giant Slalom"), "Black");
```

# --seed--

## --seed-contents--

```js
function getHillRating(drop, distance, type) {

  return drop;
}
```

# --solutions--

```js
function getHillRating(drop, distance, type) {
  let steepness = drop / distance;

  if (type === "Downhill") {
    steepness *= 1.2;
  } else if (type === "Slalom") {
    steepness *= 0.9;
  } else if (type === "Giant Slalom") {
    steepness *= 1.0;
  }

  if (steepness <= 0.1) {
    return "Green";
  } else if (steepness <= 0.25) {
    return "Blue";
  } else {
    return "Black";
  }
}
```
