---
id: 697a49e6ff50d756c9b6935d
title: "袟邪胁写邪薪薪褟 180: 蟹懈屑芯胁褨 褨谐褉懈 2026, 写械薪褜 1: 写械薪褜 胁褨写泻褉懈褌褌褟"
challengeType: 28
dashedName: challenge-180
---

# --description--

小褜芯谐芯写薪褨 锌芯褔懈薪邪褞褌褜褋褟 蟹懈屑芯胁褨 褨谐褉懈 2026 褉芯泻褍. 袧邪褋褌褍锌薪褨 17 写薪褨胁 锌褉懈薪械褋褍褌褜 胁邪屑 蟹邪胁写邪薪薪褟 蟹 泻芯写褍胁邪薪薪褟, 薪邪褌褏薪械薪薪褨 褑懈屑懈 褨谐褉邪屑懈.

袛谢褟 锌械褉褕芯谐芯 蟹邪胁写邪薪薪褟 胁邪屑 写邪褦褌褜褋褟 写胁芯谢褨褌械褉薪懈泄 泻芯写 泻褉邪褩薪懈, 褨 锌芯褌褉褨斜薪芯 锌芯胁械褉薪褍褌懈 械屑芯写蟹褨 锌褉邪锌芯褉褑褟 褑褨褦褩 泻褉邪褩薪懈.

袙懈泻芯褉懈褋褌芯胁褍泄褌械 褑械泄 褋锌懈褋芯泻:

| 袣褉邪褩薪邪                    | 袣芯写            | 袩褉邪锌芯褉械褑褜 |
| - | -: | :--: |
| Albania                    |            `"AL"` | `"馃嚘馃嚤"` |
| Andorra                    |            `"AD"` | `"馃嚘馃嚛"` |
| Argentina                  |            `"AR"` | `"馃嚘馃嚪"` |
| Armenia                    |            `"AM"` | `"馃嚘馃嚥"` |
| Australia                  |            `"AU"` | `"馃嚘馃嚭"` |
| Austria                    |            `"AT"` | `"馃嚘馃嚬"` |
| Azerbaijan                 |            `"AZ"` | `"馃嚘馃嚳"` |
| Belgium                    |            `"BE"` | `"馃嚙馃嚜"` |
| Benin                      |            `"BJ"` | `"馃嚙馃嚡"` |
| Bolivia                    |            `"BO"` | `"馃嚙馃嚧"` |
| Bosnia and Herzegovina     |            `"BA"` | `"馃嚙馃嚘"` |
| Brazil                     |            `"BR"` | `"馃嚙馃嚪"` |
| Bulgaria                   |            `"BG"` | `"馃嚙馃嚞"` |
| Canada                     |            `"CA"` | `"馃嚚馃嚘"` |
| Chile                      |            `"CL"` | `"馃嚚馃嚤"` |
| China                      |            `"CN"` | `"馃嚚馃嚦"` |
| Colombia                   |            `"CO"` | `"馃嚚馃嚧"` |
| Croatia                    |            `"HR"` | `"馃嚟馃嚪"` |
| Cyprus                     |            `"CY"` | `"馃嚚馃嚲"` |
| Czech Republic             |            `"CZ"` | `"馃嚚馃嚳"` |
| Denmark                    |            `"DK"` | `"馃嚛馃嚢"` |
| Ecuador                    |            `"EC"` | `"馃嚜馃嚚"` |
| Eritrea                    |            `"ER"` | `"馃嚜馃嚪"` |
| Estonia                    |            `"EE"` | `"馃嚜馃嚜"` |
| Finland                    |            `"FI"` | `"馃嚝馃嚠"` |
| France                     |            `"FR"` | `"馃嚝馃嚪"` |
| Georgia                    |            `"GE"` | `"馃嚞馃嚜"` |
| Germany                    |            `"DE"` | `"馃嚛馃嚜"` |
| Great Britain              |            `"GB"` | `"馃嚞馃嚙"` |
| Greece                     |            `"GR"` | `"馃嚞馃嚪"` |
| Guinea-Bissau              |            `"GW"` | `"馃嚞馃嚰"` |
| Haiti                      |            `"HT"` | `"馃嚟馃嚬"` |
| Hong Kong                  |            `"HK"` | `"馃嚟馃嚢"` |
| Hungary                    |            `"HU"` | `"馃嚟馃嚭"` |
| Iceland                    |            `"IS"` | `"馃嚠馃嚫"` |
| India                      |            `"IN"` | `"馃嚠馃嚦"` |
| Iran                       |            `"IR"` | `"馃嚠馃嚪"` |
| Ireland                    |            `"IE"` | `"馃嚠馃嚜"` |
| Israel                     |            `"IL"` | `"馃嚠馃嚤"` |
| Italy                      |            `"IT"` | `"馃嚠馃嚬"` |
| Jamaica                    |            `"JM"` | `"馃嚡馃嚥"` |
| Japan                      |            `"JP"` | `"馃嚡馃嚨"` |
| Kazakhstan                 |            `"KZ"` | `"馃嚢馃嚳"` |
| Kenya                      |            `"KE"` | `"馃嚢馃嚜"` |
| Kosovo                     |            `"XK"` | `"馃嚱馃嚢"` |
| Kyrgyzstan                 |            `"KG"` | `"馃嚢馃嚞"` |
| Latvia                     |            `"LV"` | `"馃嚤馃嚮"` |
| Lebanon                    |            `"LB"` | `"馃嚤馃嚙"` |
| Liechtenstein              |            `"LI"` | `"馃嚤馃嚠"` |
| Lithuania                  |            `"LT"` | `"馃嚤馃嚬"` |
| Luxembourg                 |            `"LU"` | `"馃嚤馃嚭"` |
| Madagascar                 |            `"MG"` | `"馃嚥馃嚞"` |
| Malaysia                   |            `"MY"` | `"馃嚥馃嚲"` |
| Malta                      |            `"MT"` | `"馃嚥馃嚬"` |
| Mexico                     |            `"MX"` | `"馃嚥馃嚱"` |
| Moldova                    |            `"MD"` | `"馃嚥馃嚛"` |
| Monaco                     |            `"MC"` | `"馃嚥馃嚚"` |
| Mongolia                   |            `"MN"` | `"馃嚥馃嚦"` |
| Montenegro                 |            `"ME"` | `"馃嚥馃嚜"` |
| Morocco                    |            `"MA"` | `"馃嚥馃嚘"` |
| Netherlands                |            `"NL"` | `"馃嚦馃嚤"` |
| New Zealand                |            `"NZ"` | `"馃嚦馃嚳"` |
| Nigeria                    |            `"NG"` | `"馃嚦馃嚞"` |
| North Macedonia            |            `"MK"` | `"馃嚥馃嚢"` |
| Norway                     |            `"NO"` | `"馃嚦馃嚧"` |
| Pakistan                   |            `"PK"` | `"馃嚨馃嚢"` |
| Philippines                |            `"PH"` | `"馃嚨馃嚟"` |
| Poland                     |            `"PL"` | `"馃嚨馃嚤"` |
| Portugal                   |            `"PT"` | `"馃嚨馃嚬"` |
| Puerto Rico                |            `"PR"` | `"馃嚨馃嚪"` |
| Romania                    |            `"RO"` | `"馃嚪馃嚧"` |
| San Marino                 |            `"SM"` | `"馃嚫馃嚥"` |
| Saudi Arabia               |            `"SA"` | `"馃嚫馃嚘"` |
| Serbia                     |            `"RS"` | `"馃嚪馃嚫"` |
| Singapore                  |            `"SG"` | `"馃嚫馃嚞"` |
| Slovakia                   |            `"SK"` | `"馃嚫馃嚢"` |
| Slovenia                   |            `"SI"` | `"馃嚫馃嚠"` |
| South Africa               |            `"ZA"` | `"馃嚳馃嚘"` |
| South Korea                |            `"KR"` | `"馃嚢馃嚪"` |
| Spain                      |            `"ES"` | `"馃嚜馃嚫"` |
| Sweden                     |            `"SE"` | `"馃嚫馃嚜"` |
| Switzerland                |            `"CH"` | `"馃嚚馃嚟"` |
| Thailand                   |            `"TH"` | `"馃嚬馃嚟"` |
| Trinidad & Tobago          |            `"TT"` | `"馃嚬馃嚬"` |
| Turkey                     |            `"TR"` | `"馃嚬馃嚪"` |
| Ukraine                    |            `"UA"` | `"馃嚭馃嚘"` |
| United Arab Emirates       |            `"AE"` | `"馃嚘馃嚜"` |
| United States              |            `"US"` | `"馃嚭馃嚫"` |
| Uruguay                    |            `"UY"` | `"馃嚭馃嚲"` |
| Uzbekistan                 |            `"UZ"` | `"馃嚭馃嚳"` |
| Venezuela                  |            `"VE"` | `"馃嚮馃嚜"` |

# --hints--

`getFlag("AL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚤"`.

```js
assert.equal(getFlag("AL"), "馃嚘馃嚤");
```

`getFlag("AD")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚛"`.

```js
assert.equal(getFlag("AD"), "馃嚘馃嚛");
```

`getFlag("AR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚪"`.

```js
assert.equal(getFlag("AR"), "馃嚘馃嚪")
```

`getFlag("AM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚥"`.

```js
assert.equal(getFlag("AM"), "馃嚘馃嚥")
```

`getFlag("AU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚭"`.

```js
assert.equal(getFlag("AU"), "馃嚘馃嚭")
```

`getFlag("AT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚬"`.

```js
assert.equal(getFlag("AT"), "馃嚘馃嚬")
```

`getFlag("AZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚳"`.

```js
assert.equal(getFlag("AZ"), "馃嚘馃嚳")
```

`getFlag("BE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚜"`.

```js
assert.equal(getFlag("BE"), "馃嚙馃嚜")
```

`getFlag("BJ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚡"`.

```js
assert.equal(getFlag("BJ"), "馃嚙馃嚡")
```

`getFlag("BO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚧"`.

```js
assert.equal(getFlag("BO"), "馃嚙馃嚧")
```

`getFlag("BA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚘"`.

```js
assert.equal(getFlag("BA"), "馃嚙馃嚘")
```

`getFlag("BR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚪"`.

```js
assert.equal(getFlag("BR"), "馃嚙馃嚪")
```

`getFlag("BG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚙馃嚞"`.

```js
assert.equal(getFlag("BG"), "馃嚙馃嚞")
```

`getFlag("CA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚘"`.

```js
assert.equal(getFlag("CA"), "馃嚚馃嚘")
```

`getFlag("CL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚤"`.

```js
assert.equal(getFlag("CL"), "馃嚚馃嚤")
```

`getFlag("CN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚦"`.

```js
assert.equal(getFlag("CN"), "馃嚚馃嚦")
```

`getFlag("CO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚧"`.

```js
assert.equal(getFlag("CO"), "馃嚚馃嚧")
```

`getFlag("HR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚪"`.

```js
assert.equal(getFlag("HR"), "馃嚟馃嚪")
```

`getFlag("CY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚲"`.

```js
assert.equal(getFlag("CY"), "馃嚚馃嚲")
```

`getFlag("CZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚳"`.

```js
assert.equal(getFlag("CZ"), "馃嚚馃嚳")
```

`getFlag("DK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚛馃嚢"`.

```js
assert.equal(getFlag("DK"), "馃嚛馃嚢")
```

`getFlag("EC")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚚"`.

```js
assert.equal(getFlag("EC"), "馃嚜馃嚚")
```

`getFlag("ER")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚪"`.

```js
assert.equal(getFlag("ER"), "馃嚜馃嚪")
```

`getFlag("EE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚜"`.

```js
assert.equal(getFlag("EE"), "馃嚜馃嚜")
```

`getFlag("FI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚝馃嚠"`.

```js
assert.equal(getFlag("FI"), "馃嚝馃嚠")
```

`getFlag("FR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚝馃嚪"`.

```js
assert.equal(getFlag("FR"), "馃嚝馃嚪")
```

`getFlag("GE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚜"`.

```js
assert.equal(getFlag("GE"), "馃嚞馃嚜")
```

`getFlag("DE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚛馃嚜"`.

```js
assert.equal(getFlag("DE"), "馃嚛馃嚜")
```

`getFlag("GB")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚙"`.

```js
assert.equal(getFlag("GB"), "馃嚞馃嚙")
```

`getFlag("GR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚪"`.

```js
assert.equal(getFlag("GR"), "馃嚞馃嚪")
```

`getFlag("GW")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚞馃嚰"`.

```js
assert.equal(getFlag("GW"), "馃嚞馃嚰")
```

`getFlag("HT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚬"`.

```js
assert.equal(getFlag("HT"), "馃嚟馃嚬")
```

`getFlag("HK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚢"`.

```js
assert.equal(getFlag("HK"), "馃嚟馃嚢")
```

`getFlag("HU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚟馃嚭"`.

```js
assert.equal(getFlag("HU"), "馃嚟馃嚭")
```

`getFlag("IS")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚫"`.

```js
assert.equal(getFlag("IS"), "馃嚠馃嚫")
```

`getFlag("IN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚦"`.

```js
assert.equal(getFlag("IN"), "馃嚠馃嚦")
```

`getFlag("IR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚪"`.

```js
assert.equal(getFlag("IR"), "馃嚠馃嚪")
```

`getFlag("IE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚜"`.

```js
assert.equal(getFlag("IE"), "馃嚠馃嚜")
```

`getFlag("IL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚤"`.

```js
assert.equal(getFlag("IL"), "馃嚠馃嚤")
```

`getFlag("IT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚠馃嚬"`.

```js
assert.equal(getFlag("IT"), "馃嚠馃嚬")
```

`getFlag("JM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚡馃嚥"`.

```js
assert.equal(getFlag("JM"), "馃嚡馃嚥")
```

`getFlag("JP")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚡馃嚨"`.

```js
assert.equal(getFlag("JP"), "馃嚡馃嚨")
```

`getFlag("KZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚳"`.

```js
assert.equal(getFlag("KZ"), "馃嚢馃嚳")
```

`getFlag("KE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚜"`.

```js
assert.equal(getFlag("KE"), "馃嚢馃嚜")
```

`getFlag("XK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚱馃嚢"`.

```js
assert.equal(getFlag("XK"), "馃嚱馃嚢")
```

`getFlag("KG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚞"`.

```js
assert.equal(getFlag("KG"), "馃嚢馃嚞")
```

`getFlag("LV")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚮"`.

```js
assert.equal(getFlag("LV"), "馃嚤馃嚮")
```

`getFlag("LB")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚙"`.

```js
assert.equal(getFlag("LB"), "馃嚤馃嚙")
```

`getFlag("LI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚠"`.

```js
assert.equal(getFlag("LI"), "馃嚤馃嚠")
```

`getFlag("LT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚬"`.

```js
assert.equal(getFlag("LT"), "馃嚤馃嚬")
```

`getFlag("LU")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚤馃嚭"`.

```js
assert.equal(getFlag("LU"), "馃嚤馃嚭")
```

`getFlag("MG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚞"`.

```js
assert.equal(getFlag("MG"), "馃嚥馃嚞")
```

`getFlag("MY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚲"`.

```js
assert.equal(getFlag("MY"), "馃嚥馃嚲")
```

`getFlag("MT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚬"`.

```js
assert.equal(getFlag("MT"), "馃嚥馃嚬")
```

`getFlag("MX")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚱"`.

```js
assert.equal(getFlag("MX"), "馃嚥馃嚱")
```

`getFlag("MD")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚛"`.

```js
assert.equal(getFlag("MD"), "馃嚥馃嚛")
```

`getFlag("MC")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚚"`.

```js
assert.equal(getFlag("MC"), "馃嚥馃嚚")
```

`getFlag("MN")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚦"`.

```js
assert.equal(getFlag("MN"), "馃嚥馃嚦")
```

`getFlag("ME")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚜"`.

```js
assert.equal(getFlag("ME"), "馃嚥馃嚜")
```

`getFlag("MA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚘"`.

```js
assert.equal(getFlag("MA"), "馃嚥馃嚘")
```

`getFlag("NL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚤"`.

```js
assert.equal(getFlag("NL"), "馃嚦馃嚤")
```

`getFlag("NZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚳"`.

```js
assert.equal(getFlag("NZ"), "馃嚦馃嚳")
```

`getFlag("NG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚞"`.

```js
assert.equal(getFlag("NG"), "馃嚦馃嚞")
```

`getFlag("MK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚥馃嚢"`.

```js
assert.equal(getFlag("MK"), "馃嚥馃嚢")
```

`getFlag("NO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚦馃嚧"`.

```js
assert.equal(getFlag("NO"), "馃嚦馃嚧")
```

`getFlag("PK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚢"`.

```js
assert.equal(getFlag("PK"), "馃嚨馃嚢")
```

`getFlag("PH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚟"`.

```js
assert.equal(getFlag("PH"), "馃嚨馃嚟")
```

`getFlag("PL")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚤"`.

```js
assert.equal(getFlag("PL"), "馃嚨馃嚤")
```

`getFlag("PT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚬"`.

```js
assert.equal(getFlag("PT"), "馃嚨馃嚬")
```

`getFlag("PR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚨馃嚪"`.

```js
assert.equal(getFlag("PR"), "馃嚨馃嚪")
```

`getFlag("RO")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚪馃嚧"`.

```js
assert.equal(getFlag("RO"), "馃嚪馃嚧")
```

`getFlag("SM")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚥"`.

```js
assert.equal(getFlag("SM"), "馃嚫馃嚥")
```

`getFlag("SA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚘"`.

```js
assert.equal(getFlag("SA"), "馃嚫馃嚘")
```

`getFlag("RS")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚪馃嚫"`.

```js
assert.equal(getFlag("RS"), "馃嚪馃嚫")
```

`getFlag("SG")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚞"`.

```js
assert.equal(getFlag("SG"), "馃嚫馃嚞")
```

`getFlag("SK")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚢"`.

```js
assert.equal(getFlag("SK"), "馃嚫馃嚢")
```

`getFlag("SI")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚠"`.

```js
assert.equal(getFlag("SI"), "馃嚫馃嚠")
```

`getFlag("ZA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚳馃嚘"`.

```js
assert.equal(getFlag("ZA"), "馃嚳馃嚘")
```

`getFlag("KR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚢馃嚪"`.

```js
assert.equal(getFlag("KR"), "馃嚢馃嚪")
```

`getFlag("ES")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚜馃嚫"`.

```js
assert.equal(getFlag("ES"), "馃嚜馃嚫")
```

`getFlag("SE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚫馃嚜"`.

```js
assert.equal(getFlag("SE"), "馃嚫馃嚜")
```

`getFlag("CH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚚馃嚟"`.

```js
assert.equal(getFlag("CH"), "馃嚚馃嚟")
```

`getFlag("TH")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚟"`.

```js
assert.equal(getFlag("TH"), "馃嚬馃嚟")
```

`getFlag("TT")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚬"`.

```js
assert.equal(getFlag("TT"), "馃嚬馃嚬")
```

`getFlag("TR")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚬馃嚪"`.

```js
assert.equal(getFlag("TR"), "馃嚬馃嚪")
```

`getFlag("UA")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚘"`.

```js
assert.equal(getFlag("UA"), "馃嚭馃嚘")
```

`getFlag("AE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚘馃嚜"`.

```js
assert.equal(getFlag("AE"), "馃嚘馃嚜")
```

`getFlag("US")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚫"`.

```js
assert.equal(getFlag("US"), "馃嚭馃嚫")
```

`getFlag("UY")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚲"`.

```js
assert.equal(getFlag("UY"), "馃嚭馃嚲")
```

`getFlag("UZ")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚭馃嚳"`.

```js
assert.equal(getFlag("UZ"), "馃嚭馃嚳")
```

`getFlag("VE")` 屑邪褦 锌芯胁械褉褌邪褌懈 `"馃嚮馃嚜"`.

```js
assert.equal(getFlag("VE"), "馃嚮馃嚜")
```

# --seed--

## --seed-contents--

```js
function getFlag(code) {

  return code;
}
```

# --solutions--

```js
function getFlag(code) {
  return [...code]
    .map(char => String.fromCodePoint(127397 + char.charCodeAt()))
    .join("");
}
```
