---
id: 697a49e9860d24853adef67e
title: "Завдання 193: зимові ігри 2026, день 14: лижний альпінізм"
challengeType: 28
dashedName: challenge-193
---

# --description--

Враховуючи глибину снігу та нахил гори, визначте, чи існує ризик лавини.

- Значення глибини снігу — це `"Shallow"`, `"Moderate"` або `"Deep"`.
- Значення нахилу — це `"Gentle"`, `"Steep"` або `"Very Steep"`.

Поверніть `"Safe"` або `"Risky"` відповідно до цієї таблиці:

| |`"Shallow"`|`"Moderate"`|`"Deep"`|
|-|-|-|-|
|`"Gentle"`|`"Safe"`|`"Safe"`|`"Safe"`|
|`"Steep"`|`"Safe"`|`"Risky"`|`"Risky"`|
|`"Very Steep"`|`"Safe"`| `"Risky"`|`"Risky"`|

# --hints--

`avalancheRisk("Shallow", "Gentle")` має повертати `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Gentle"), "Safe");
```

`avalancheRisk("Shallow", "Steep")` має повертати `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Steep"), "Safe");
```

`avalancheRisk("Shallow", "Very Steep")` має повертати `"Safe"`.

```js
assert.equal(avalancheRisk("Shallow", "Very Steep"), "Safe");
```

`avalancheRisk("Moderate", "Gentle")` має повертати `"Safe"`.

```js
assert.equal(avalancheRisk("Moderate", "Gentle"), "Safe");
```

`avalancheRisk("Moderate", "Steep")` має повертати `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Steep"), "Risky");
```

`avalancheRisk("Moderate", "Very Steep")` має повертати `"Risky"`.

```js
assert.equal(avalancheRisk("Moderate", "Very Steep"), "Risky");
```

`avalancheRisk("Deep", "Gentle")` має повертати `"Safe"`.

```js
assert.equal(avalancheRisk("Deep", "Gentle"), "Safe");
```

`avalancheRisk("Deep", "Steep")` має повертати `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Steep"), "Risky");
```

`avalancheRisk("Deep", "Very Steep")` має повертати `"Risky"`.

```js
assert.equal(avalancheRisk("Deep", "Very Steep"), "Risky");
```

# --seed--

## --seed-contents--

```js
function avalancheRisk(snowDepth, slope) {

  return snowDepth;
}
```

# --solutions--

```js
function avalancheRisk(snowDepth, slope) {
  const riskTable = {
    Gentle: {
      Shallow: "Safe",
      Moderate: "Safe",
      Deep: "Safe",
    },
    Steep: {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
    "Very Steep": {
      Shallow: "Safe",
      Moderate: "Risky",
      Deep: "Risky",
    },
  };

  return riskTable[slope][snowDepth];
}
```
