---
id: 697a49e9860d24853adef67d
title: "Завдання 192: зимові ігри 2026, день 13: лижне двоборство"
challengeType: 28
dashedName: challenge-192
---

# --description--

Дано масив балів за стрибки спортсменів, обчисліть час затримки їхнього старту для крос-кантрі частини лижного двоборства.

Спортсмен із найвищим балом за стрибок стартує першим (затримка 0 секунд). Усі інші спортсмени стартують пізніше залежно від того, наскільки їхній бал за стрибок відстає від найкращого.

Щоб обчислити затримку для кожного спортсмена, відніміть бал спортсмена за стрибок від найкращого загального балу та помножте результат на 1.5. Округліть затримку вгору до найближчого цілого числа.

# --hints--

`calculateStartDelays([120, 110, 125])` має повертати `[8, 23, 0]`.

```js
assert.deepEqual(calculateStartDelays([120, 110, 125]), [8, 23, 0]);
```

`calculateStartDelays([118, 125, 122, 120])` має повертати `[11, 0, 5, 8]`.

```js
assert.deepEqual(calculateStartDelays([118, 125, 122, 120]), [11, 0, 5, 8]);
```

`calculateStartDelays([100, 105, 95, 110, 120, 115, 108])` має повертати `[30, 23, 38, 15, 0, 8, 18]`.

```js
assert.deepEqual(calculateStartDelays([100, 105, 95, 110, 120, 115, 108]), [30, 23, 38, 15, 0, 8, 18]);
```

`calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124])` має повертати `[3, 11, 6, 18, 21, 15, 8, 26, 0, 12]`.

```js
assert.deepEqual(calculateStartDelays([130, 125, 128, 120, 118, 122, 127, 115, 132, 124]), [3, 11, 6, 18, 21, 15, 8, 26, 0, 12]);
```

# --seed--

## --seed-contents--

```js
function calculateStartDelays(jumpScores) {

  return jumpScores
}
```

# --solutions--

```js
function calculateStartDelays(jumpScores) {
  const bestJump = Math.max(...jumpScores);
  return jumpScores.map(score => Math.ceil((bestJump - score) * 1.5));
}
```
