---
id: 697a49e6ff50d756c9b69365
title: "Завдання 188: зимові ігри 2026, день 9: скелетон"
challengeType: 28
dashedName: challenge-188
---

# --description--

Дано рядок, що представляє повороти на трасі для скелетону. Визначте складність траси.

- Даний рядок міститиме лише літери:
  - `"L"` для лівого повороту
  - `"R"` для правого повороту
  - `"S"` для прямого відрізка

- Кожна зміна напрямку додає 15 балів (`"L"`, за яким іде `"R"` або навпаки).
- Всі інші повороти додають по 5 балів кожен (усі інші символи `"L"` або `"R"`).
- Прямі відрізки додають 0 балів.

Складність траси базується на загальній сумі балів. Поверніть:

- `"Easy"`, якщо сума від 0 до 100
- `"Medium"`, якщо сума від 101 до 200
- `"Hard"`, якщо сума понад 200

# --hints--

`getDifficulty("SLSLLSRRLSRLRL")` має повернути `"Easy"`.

```js
assert.equal(getDifficulty("SLSLLSRRLSRLRL"), "Easy");
```

`getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS")` має повернути `"Hard"`.

```js
assert.equal(getDifficulty("LLRSLRLRSLLRLRSLRRLRSRLLS"), "Hard");
```

`getDifficulty("SRRRRLSLLRLRSSRLSRL")` має повернути `"Medium"`.

```js
assert.equal(getDifficulty("SRRRRLSLLRLRSSRLSRL"), "Medium");
```

`getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL")` має повернути `"Hard"`.

```js
assert.equal(getDifficulty("LSRLRLSRLRLSLRSLRLLRLSRLRLRSL"), "Hard");
```

`getDifficulty("SLLSSLRLSLSLRSLSSLRL")` має повернути `"Medium"`.

```js
assert.equal(getDifficulty("SLLSSLRLSLSLRSLSSLRL"), "Medium");
```

`getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR")` має повернути `"Easy"`.

```js
assert.equal(getDifficulty("SRSLSRSLSRRSLSRSRSLSRLSRSR"), "Easy");
```

# --seed--

## --seed-contents--

```js
function getDifficulty(track) {

  return track;
}
```

# --solutions--

```js
function getDifficulty(track) {
  let score = 0;

  for (let i = 0; i < track.length; i++) {
    const current = track[i];
    const previous = i > 0 ? track[i - 1] : null;

    if (current === 'S') {
      score += 0;
    } else if (current === 'L' || current === 'R') {
      if (previous && previous !== 'S' && previous !== current) {
        score += 15;
      } else if (previous !== 'S') {
        score += 5;
      } else {
        score += 5;
      }
    }
  }

  if (score <= 100) return 'Easy';
  if (score <= 200) return 'Medium';
  return 'Hard';
}
```
