---
id: 697a49e6ff50d756c9b6935e
title: "Завдання 181: зимові ігри 2026, день 2: сноубординг"
challengeType: 28
dashedName: challenge-181
---

# --description--

Визначте позу сноубордиста при приземленні, якщо відомі його початкова поза та обертання в градусах.

- Поза сноубордиста може бути або `"Regular"`, або `"Goofy"`.
- Обертання трюку — це кратні 90 градусів. Позитивне значення означає обертання за годинниковою стрілкою, від’ємне — проти годинникової стрілки.
- Поза при приземленні змінюється на протилежну кожні 180 градусів обертання.

Наприклад, якщо задано `"Regular"` і `90`, поверніть `"Regular"`. Якщо задано `"Regular"` і `180` градусів, поверніть `"Goofy"`.

# --hints--

`getLandingStance("Regular", 90)` має повернути `"Regular"`.

```js
assert.equal(getLandingStance("Regular", 90), "Regular");
```

`getLandingStance("Regular", 180)` має повернути `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 180), "Goofy");
```

`getLandingStance("Goofy", -270)` має повернути `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -270), "Regular");
```

`getLandingStance("Regular", 2340)` має повернути `"Goofy"`.

```js
assert.equal(getLandingStance("Regular", 2340), "Goofy");
```

`getLandingStance("Goofy", 2160)` має повернути `"Goofy"`.

```js
assert.equal(getLandingStance("Goofy", 2160), "Goofy");
```

`getLandingStance("Goofy", -540)` має повернути `"Regular"`.

```js
assert.equal(getLandingStance("Goofy", -540), "Regular");
```

# --seed--

## --seed-contents--

```js
function getLandingStance(startStance, rotation) {

  return startStance;
}
```

# --solutions--

```js
function getLandingStance(startStance, rotation) {
  const flips = Math.floor(Math.abs(rotation) / 180);
  if (flips % 2 === 0) {
    return startStance;
  } else {
    return startStance === "Regular" ? "Goofy" : "Regular";
  }
}
```
