---
id: 697a49e9860d24853adef67b
title: "Завдання 190: зимові ігри 2026, день 11: хокей на льоду"
challengeType: 28
dashedName: challenge-190
---

# --description--

Дано масив із 6 команд з хокею на льоду та їхніми результатами після кругового турніру. Визначте пари для півфінального раунду.

- Кожен елемент масиву містить команду та її результати у форматі `"TEAM: W-OTW-OTL-L"`. Де:
  - `"W"` — кількість перемог у основний час, кожна оцінюється в 3 очки
  - `"OTW"` — кількість перемог у додатковий час, кожна оцінюється в 2 очки
  - `"OTL"` — кількість поразок у додатковий час, кожна оцінюється в 1 очко
  - `"L"` — кількість поразок у основний час, кожна оцінюється в 0 очок
  
Наприклад, `"FIN: 2-2-1-0"` матиме 11 очок після підсумовування результатів.
  
Знайдіть загальну кількість очок для кожної команди і поверніть `"The semi-final games will be (1st) vs (4th) and (2nd) vs (3rd)."`. Наприклад, `"The semi-final games will be FIN vs SWE and CAN vs USA."`

# --hints--

`getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"])` має повернути `"The semi-final games will be FIN vs SWE and CAN vs USA."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-2-0-1", "FIN: 2-2-1-0", "GER: 1-0-1-3", "SUI: 0-1-3-1", "SWE: 1-1-2-1", "USA: 2-1-0-2"]), "The semi-final games will be FIN vs SWE and CAN vs USA.");
```

`getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"])` має повернути `"The semi-final games will be USA vs CZE and CAN vs FIN."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 2-1-1-1", "CZE: 1-1-1-2", "FIN: 1-2-1-1", "NOR: 0-1-1-3", "SLO: 1-0-1-3", "USA: 5-0-0-0"]), "The semi-final games will be USA vs CZE and CAN vs FIN.");
```

`getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"])` має повернути `"The semi-final games will be CAN vs ITA and USA vs CZE."`.

```js
assert.equal(getSemifinalMatchups(["CAN: 3-2-0-0", "CZE: 2-1-2-0", "LAT: 0-0-1-4", "ITA: 1-1-1-2", "DEN: 1-0-0-4", "USA: 3-1-1-0"]), "The semi-final games will be CAN vs ITA and USA vs CZE.");
```

`getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"])` має повернути `"The semi-final games will be JPN vs ITA and AUT vs KOR."`.

```js
assert.equal(getSemifinalMatchups(["AUT: 2-2-1-0", "DEN: 1-0-0-4", "ITA: 1-1-1-2", "JPN: 3-2-0-0", "KOR: 2-1-2-0", "LAT: 0-0-1-4"]), "The semi-final games will be JPN vs ITA and AUT vs KOR.");
```

# --seed--

## --seed-contents--

```js
function getSemifinalMatchups(teams) {

  return teams;
}
```

# --solutions--

```js
function getSemifinalMatchups(teams) {
    const teamPoints = teams.map(teamStr => {
        const [name, record] = teamStr.split(": ");
        const [W, OTW, OTL, L] = record.split("-").map(Number);
        const points = W * 3 + OTW * 2 + OTL * 1 + L * 0;
        return { name, points };
    });

    teamPoints.sort((a, b) => b.points - a.points);
    const [first, second, third, fourth] = teamPoints;

    return `The semi-final games will be ${first.name} vs ${fourth.name} and ${second.name} vs ${third.name}.`;
}
```
