---
id: 698a1a73ade5ac0e19180fa3
title: "Завдання 200: підрахунок літер і чисел"
challengeType: 28
dashedName: challenge-200
---

# --description--

Дано рядок, поверніть повідомлення з підрахунком, скільки в ньому міститься літер і чисел.

- Літери — це `A-Z` і `a-z`.
- Числа — це `0-9`.
- Ігноруйте всі інші символи.

Поверніть `"The string has X letters and Y numbers."`, де `"X"` — це кількість літер, а `"Y"` — кількість чисел. Якщо кількість будь-якого елемента дорівнює 1, використовуйте однину для цього елемента. Наприклад: `"1 letter"` замість `"1 letters"` і `"1 number"` замість `"1 numbers"`.

# --hints--

`countLettersAndNumbers("helloworld123")` має повертати `"The string has 10 letters and 3 numbers."`.

```js
assert.equal(countLettersAndNumbers("helloworld123"), "The string has 10 letters and 3 numbers.");
```

`countLettersAndNumbers("Catch 22")` має повертати `"The string has 5 letters and 2 numbers."`.

```js
assert.equal(countLettersAndNumbers("Catch 22"), "The string has 5 letters and 2 numbers.");
```

`countLettersAndNumbers("A1!")` має повертати `"The string has 1 letter and 1 number."`.

```js
assert.equal(countLettersAndNumbers("A1!"), "The string has 1 letter and 1 number.");
```

`countLettersAndNumbers("12345")` має повертати `"The string has 0 letters and 5 numbers."`.

```js
assert.equal(countLettersAndNumbers("12345"), "The string has 0 letters and 5 numbers.");
```

`countLettersAndNumbers("password")` має повертати `"The string has 8 letters and 0 numbers."`.

```js
assert.equal(countLettersAndNumbers("password"), "The string has 8 letters and 0 numbers.");
```

# --seed--

## --seed-contents--

```js
function countLettersAndNumbers(str) {

  return str;
}
```

# --solutions--

```js
function countLettersAndNumbers(str) {
  let letterCount = 0;
  let numberCount = 0;

  for (let char of str) {
    if ((char >= 'A' && char <= 'Z') || (char >= 'a' && char <= 'z')) {
      letterCount++;
    } else if (char >= '0' && char <= '9') {
      numberCount++;
    }
  }

  const letterWord = letterCount === 1 ? "letter" : "letters";
  const numberWord = numberCount === 1 ? "number" : "numbers";

  return `The string has ${letterCount} ${letterWord} and ${numberCount} ${numberWord}.`;
}
```
