---
id: 698a1a73ade5ac0e19180fa8
title: "Завдання 205: підрахунок ідеальних кубів"
challengeType: 28
dashedName: challenge-205
---

# --description--

Дано два цілі числа, визначте, скільки ідеальних кубів існує в діапазоні між цими двома числами включно.

- Нижче число не обов’язково є першим аргументом.
- Число є ідеальним кубом, якщо існує ціле число (`n`), для якого `n * n * n = number`. Наприклад, 27 — це ідеальний куб, тому що `3 * 3 * 3 = 27`.

# --hints--

`countPerfectCubes(3, 30)` має повертати `2`.

```js
assert.equal(countPerfectCubes(3, 30), 2);
```

`countPerfectCubes(1, 30)` має повертати `3`.

```js
assert.equal(countPerfectCubes(1, 30), 3);
```

`countPerfectCubes(30, 0)` має повертати `4`.

```js
assert.equal(countPerfectCubes(30, 0), 4);
```

`countPerfectCubes(-64, 64)` має повертати `9`.

```js
assert.equal(countPerfectCubes(-64, 64), 9);
```

`countPerfectCubes(9214, -8127)` має повертати `41`.

```js
assert.equal(countPerfectCubes(9214, -8127), 41);
```

# --seed--

## --seed-contents--

```js
function countPerfectCubes(a, b) {

  return a;
}
```

# --solutions--

```js
function countPerfectCubes(a, b) {
  const start = Math.min(a, b);
  const end = Math.max(a, b);

  let count = 0;
  let n = 0;

  while (n ** 3 <= end) {
    if (n ** 3 >= start) count++;
    n++;
  }

  n = -1;
  while (n ** 3 >= start) {
    if (n ** 3 <= end) count++;
    n--;
  }

  return count;
}
```
