---
id: 698a1a73ade5ac0e19180fa1
title: "Завдання 198: підрахунок робочих днів"
challengeType: 28
dashedName: challenge-198
---

# --description--

За заданою датою початку та датою кінця поверніть число робочих днів між ними.

- Задані дати у форматі `"YYYY-MM-DD"`.
- Робочими днями вважаються будні (з понеділка по п’ятницю).
- Вихідні дні (субота та неділя) не є робочими.
- Підрахунок включає як дату початку, так і дату кінця.

# --hints--

`countBusinessDays("2026-02-24", "2026-02-26")` має повертати `3`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-26"), 3);
```

`countBusinessDays("2026-02-24", "2026-02-28")` має повертати `4`.

```js
assert.equal(countBusinessDays("2026-02-24", "2026-02-28"), 4);
```

`countBusinessDays("2026-02-21", "2026-03-01")` має повертати `5`.

```js
assert.equal(countBusinessDays("2026-02-21", "2026-03-01"), 5);
```

`countBusinessDays("2026-03-08", "2026-03-17")` має повертати `7`.

```js
assert.equal(countBusinessDays("2026-03-08", "2026-03-17"), 7);
```

`countBusinessDays("2026-02-24", "2027-02-24")` має повертати `262`.

```js
assert.equal(countBusinessDays("2026-02-24", "2027-02-24"), 262);
```

# --seed--

## --seed-contents--

```js
function countBusinessDays(start, end) {

  return start;
}
```

# --solutions--

```js
function countBusinessDays(start, end) {
  const startDate = new Date(start + "T00:00:00Z");
  const endDate = new Date(end + "T00:00:00Z");

  let count = 0;
  let current = new Date(startDate);

  while (current <= endDate) {
    const day = current.getUTCDay();
    if (day !== 0 && day !== 6) count++;
    current.setUTCDate(current.getUTCDate() + 1);
  }

  return count;
}
```
