---
id: 698a1a73ade5ac0e19180fa0
title: "Завдання 197: сумісність груп крові"
challengeType: 28
dashedName: challenge-197
---

# --description--

За заданими групою крові донора та групою крові реципієнта визначте, чи може донор здати кров реципієнту.

Кожна група крові складається з:

- Літери: `"A"`, `"B"`, `"AB"` або `"O"`
- І фактора Rh: `"+"` або `"-"`

Групи крові будуть однією з дійсних літер, за якою слідує фактор Rh. Наприклад, `"AB+"` і `"O-"` — це дійсні групи крові.

Правила для літер:

- `"O"` може здавати кров іншим типам літер.
- `"A"` може здавати кров `"A"` і `"AB"`.
- `"B"` може здавати кров `"B"` і `"AB"`.
- `"AB"` може здавати кров лише `"AB"`.

Правила для Rh:

- Від’ємний (`"-"`) може здавати кров як `"-"`, так і `"+"`.
- Позитивний (`"+"`) може здавати кров лише `"+"`.

Для того, щоб донор міг здати кров реципієнту, мають виконуватися обидва правила — і для літери, і для Rh.

# --hints--

`canDonate("B+", "B+")` має повертати `true`.

```js
assert.isTrue(canDonate("B+", "B+"));
```

`canDonate("O-", "AB-")` має повертати `true`.

```js
assert.isTrue(canDonate("O-", "AB-"));
```

`canDonate("O+", "A-")` має повертати `false`.

```js
assert.isFalse(canDonate("O+", "A-"));
```

`canDonate("A+", "AB+")` має повертати `true`.

```js
assert.isTrue(canDonate("A+", "AB+"));
```

`canDonate("A-", "B-")` має повертати `false`.

```js
assert.isFalse(canDonate("A-", "B-"));
```

`canDonate("B-", "AB+")` має повертати `true`.

```js
assert.isTrue(canDonate("B-", "AB+"));
```

`canDonate("B-", "A+")` має повертати `false`.

```js
assert.isFalse(canDonate("B-", "A+"));
```

`canDonate("O-", "O+")` має повертати `true`.

```js
assert.isTrue(canDonate("O-", "O+"));
```

`canDonate("O+", "O-")` має повертати `false`.

```js
assert.isFalse(canDonate("O+", "O-"));
```

`canDonate("AB+", "AB-")` має повертати `false`.

```js
assert.isFalse(canDonate("AB+", "AB-"));
```

# --seed--

## --seed-contents--

```js
function canDonate(donor, recipient) {

  return donor;
}
```

# --solutions--

```js
function canDonate(donor, recipient) {
  const donorType = donor.slice(0, -1);
  const donorRh = donor.slice(-1);

  const recipientType = recipient.slice(0, -1);
  const recipientRh = recipient.slice(-1);

  const aboCompatibility = {
    "O": ["A", "B", "AB", "O"],
    "A": ["A", "AB"],
    "B": ["B", "AB"],
    "AB": ["AB"]
  };

  const aboMatch = aboCompatibility[donorType].includes(recipientType);
  const rhMatch =
    donorRh === "-" || (donorRh === "+" && recipientRh === "+");

  return aboMatch && rhMatch;
}
```
