---
id: 698a1a863194f1f4e63f6460
title: "Завдання 209: розмір елемента"
challengeType: 28
dashedName: challenge-209
---

# --description--

За заданим розміром вікна, шириною елемента у одиницях ширини вікна перегляду `"vw"` та висотою елемента у одиницях висоти вікна перегляду `"vh"`, визначте розмір елемента у пікселях.

- Заданий розмір вікна та повернутий розмір елемента — це рядки у форматі `"width x height"`, `"1200 x 800"`, наприклад.

- `"vw"` одиниці — це відсоток від ширини вікна. `"50vw"`, наприклад, означає 50% ширини вікна.  
- `"vh"` одиниці — це відсоток від висоти вікна. `"50vh"`, наприклад, означає 50% висоти вікна.

# --hints--

`getElementSize("1200 x 800", "50vw", "50vh")` має повертати `"600 x 400"`.

```js
assert.equal(getElementSize("1200 x 800", "50vw", "50vh"), "600 x 400");
```

`getElementSize("320 x 480", "25vw", "50vh")` має повертати `"80 x 240"`.

```js
assert.equal(getElementSize("320 x 480", "25vw", "50vh"), "80 x 240");
```

`getElementSize("1000 x 500", "7vw", "3vh")` має повертати `"70 x 15"`.

```js
assert.equal(getElementSize("1000 x 500", "7vw", "3vh"), "70 x 15");
```

`getElementSize("1920 x 1080", "95vw", "100vh")` має повертати `"1824 x 1080"`.

```js
assert.equal(getElementSize("1920 x 1080", "95vw", "100vh"), "1824 x 1080");
```

`getElementSize("1200 x 800", "0vw", "0vh")` має повертати `"0 x 0"`.

```js
assert.equal(getElementSize("1200 x 800", "0vw", "0vh"), "0 x 0");
```

`getElementSize("1440 x 900", "100vw", "114vh")` має повертати `"1440 x 1026"`.

```js
assert.equal(getElementSize("1440 x 900", "100vw", "114vh"), "1440 x 1026");
```

# --seed--

## --seed-contents--

```js
function getElementSize(windowSize, elementVw, elementVh) {

  return windowSize;
}
```

# --solutions--

```js
function getElementSize(windowSize, elementVw, elementVh) {
  const [windowWidth, windowHeight] = windowSize
    .split(" x ")
    .map(Number);

  const vw = Number(elementVw.replace("vw", ""));
  const vh = Number(elementVh.replace("vh", ""));

  const widthPx = (vw / 100) * windowWidth;
  const heightPx = (vh / 100) * windowHeight;

  return `${widthPx} x ${heightPx}`;
}
```
