---
id: 698a1a73ade5ac0e19180fa5
title: "Завдання 202: додати розділові знаки"
challengeType: 28
dashedName: challenge-202
---

# --description--

Дано рядок речень із пропущеними крапками. Додайте крапку (`"."`) у таких місцях:

- Перед кожним пробілом, який безпосередньо передує літері у верхньому регістрі
- А також у кінці рядка

Поверніть отриманий рядок.

# --hints--

`addPunctuation("Hello world")` має повертати `"Hello world."`

```js
assert.equal(addPunctuation("Hello world"), "Hello world.");
```

`addPunctuation("Hello world It's nice today")` має повертати `"Hello world. It's nice today."`

```js
assert.equal(addPunctuation("Hello world It's nice today"), "Hello world. It's nice today.");
```

`addPunctuation("JavaScript is great Sometimes")` має повертати `"JavaScript is great. Sometimes."`

```js
assert.equal(addPunctuation("JavaScript is great Sometimes"), "JavaScript is great. Sometimes.");
```

`addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z")` має повертати `"A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z."`

```js
assert.equal(addPunctuation("A b c D e F g h I J k L m n o P Q r S t U v w X Y Z"), "A b c. D e. F g h. I. J k. L m n o. P. Q r. S t. U v w. X. Y. Z.");
```

`addPunctuation("Wait.. For it")` має повертати `"Wait... For it."`

```js
assert.equal(addPunctuation("Wait.. For it"), "Wait... For it.");
```

# --seed--

## --seed-contents--

```js
function addPunctuation(sentences) {

  return sentences;
}
```

# --solutions--

```js
function addPunctuation(sentences) {
  let result = "";

  for (let i = 0; i < sentences.length; i++) {
    if (sentences[i] === " " && i + 1 < sentences.length && /[A-Z]/.test(sentences[i + 1])) {
      result += ".";
    }
    result += sentences[i];
  }
  result += ".";

  return result;
}
```
