---
id: 698a1a863194f1f4e63f645e
title: "Завдання 207: найменший горизонтальний розрив"
challengeType: 28
dashedName: challenge-207
---

# --description--

Дано рядок, поверніть підрядок між двома ідентичними символами, між якими найменша кількість символів (найменший горизонтальний розрив).

- Завжди буде принаймні одна пара однакових символів.
- Повернутий підрядок не повинен включати однакові символи.
- Якщо два або більше горизонтальних розривів однакової довжини, поверніть символи з першого.

Наприклад, для `"ABCDAC"` поверніть `"DA"`, тому що:

- Лише `"A"` і `"C"` повторюються в рядку.
- Кількість символів між двома `"A"` символами дорівнює 3, а між `"C"` символами — 2.
- Отже, поверніть рядок між двома `"C"` символами.

# --hints--

`smallestGap("ABCDAC")` має повернути `"DA"`.

```js
assert.equal(smallestGap("ABCDAC"), "DA");
```

`smallestGap("racecar")` має повернути `"e"`.

```js
assert.equal(smallestGap("racecar"), "e");
```

`smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4")` має повернути `"#q6e&rkf(p"`.

```js
assert.equal(smallestGap("A{5e^SD*F4i!o#q6e&rkf(po8|we9+kr-2!3}=4"), "#q6e&rkf(p");
```

`smallestGap("Hello World")` має повернути `""`.

```js
assert.equal(smallestGap("Hello World"), "");
```

`smallestGap("The quick brown fox jumps over the lazy dog.")` має повернути `"fox"`.

```js
assert.equal(smallestGap("The quick brown fox jumps over the lazy dog."), "fox");
```

# --seed--

## --seed-contents--

```js
function smallestGap(str) {

  return str;
}
```

# --solutions--

```js
function smallestGap(str) {
  let minGap = Infinity;
  let result = "";

  for (let i = 0; i < str.length; i++) {
    for (let j = i + 1; j < str.length; j++) {
      if (str[i] === str[j]) {
        const gap = j - i - 1;

        if (gap < minGap) {
          minGap = gap;
          result = str.slice(i + 1, j);
        }

        break;
      }
    }
  }

  return result;
}
```
