---
id: 698a1a73ade5ac0e19180fa4
title: "Завдання 201: зсув матриці"
challengeType: 28
dashedName: challenge-201
---

# --description--

Дано матрицю (масив масивів) чисел і ціле число. Зсунути всі значення в матриці на задану величину.

- Позитивний зсув переміщує значення вправо.
- Від’ємний зсув переміщує значення вліво.

Значення повинні обертатися:

- Розглядайте матрицю як одну безперервну послідовність значень.
- Коли значення виходить за межі рядка, воно продовжує рухатися з початку наступного рядка.
- Коли значення виходить за останню позицію в матриці, воно обертається до першої позиції.
- Те саме застосовується у зворотному напрямку при зсуві вліво.

Наприклад, для:

```js
[
  [1, 2, 3],
  [4, 5, 6]
]
```

зі зсувом `1` перемістіть усі числа вправо на один:

```js
[
  [6, 1, 2],
  [3, 4, 5]
]
```

# --hints--

`shiftMatrix([[1, 2, 3], [4, 5, 6]], 1)` має повернути `[[6, 1, 2], [3, 4, 5]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], 1), [[6, 1, 2], [3, 4, 5]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6]], -1)` має повернути `[[2, 3, 4], [5, 6, 1]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6]], -1), [[2, 3, 4], [5, 6, 1]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5)` має повернути `[[5, 6, 7], [8, 9, 1], [2, 3, 4]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], 5), [[5, 6, 7], [8, 9, 1], [2, 3, 4]]);
```

`shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6)` має повернути `[[7, 8, 9], [1, 2, 3], [4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3], [4, 5, 6], [7, 8, 9]], -6), [[7, 8, 9], [1, 2, 3], [4, 5, 6]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7)` має повернути `[[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], 7), [[10, 11, 12, 13], [14, 15, 16, 1], [2, 3, 4, 5], [6, 7, 8, 9]]);
```

`shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54)` має повернути `[[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]`.

```js
assert.deepEqual(shiftMatrix([[1, 2, 3, 4], [5, 6, 7, 8], [9, 10, 11, 12], [13, 14, 15, 16]], -54), [[7, 8, 9, 10], [11, 12, 13, 14], [15, 16, 1, 2], [3, 4, 5, 6]]);
```

# --seed--

## --seed-contents--

```js
function shiftMatrix(matrix, shift) {

  return matrix;
}
```

# --solutions--

```js
function shiftMatrix(matrix, shift) {
  const rows = matrix.length;
  const cols = matrix[0].length;
  const flat = matrix.flat();
  const n = flat.length;
  const s = ((shift % n) + n) % n;
  const shifted = flat.slice(-s).concat(flat.slice(0, n - s));
  const result = [];
  for (let i = 0; i < rows; i++) {
    result.push(shifted.slice(i * cols, (i + 1) * cols));
  }

  return result;
}
```
