---
id: 698a1a73ade5ac0e19180fa9
title: "Завдання 206: значення ігрових карт"
challengeType: 28
dashedName: challenge-206
---

# --description--

Дано масив ігрових карт, поверніть новий масив із числовим значенням кожної карти.

Значення карт:

- Туз (`"A"`) має значення 1.
- Картки з числами (`"2"` - `"10"`) мають своє номінальне значення: відповідно від 2 до 10.
- Фігурні карти: Валет (`"J"`), Дама (`"Q"`) і Король (`"K"`) оцінюються в 10.

Масти:

- Кожна карта має масть: піки (`"S"`), трефи (`"C"`), бубни (`"D"`) або черви (`"H"`).

Формат карти:

- Кожна карта представлена у вигляді рядка: `"valueSuit"`. Наприклад: `"AS"` — це туз пік, `"10H"` — десять черв, а `"QC"` — дама треф.

# --hints--

`cardValues(["3H", "4D", "5S"])` має повернути `[3, 4, 5]`.

```js
assert.deepEqual(cardValues(["3H", "4D", "5S"]), [3, 4, 5]);
```

`cardValues(["AS", "10S", "10H", "6D", "7D"])` має повернути `[1, 10, 10, 6, 7]`.

```js
assert.deepEqual(cardValues(["AS", "10S", "10H", "6D", "7D"]), [1, 10, 10, 6, 7]);
```

`cardValues(["8D", "QS", "2H", "JC", "9C"])` має повернути `[8, 10, 2, 10, 9]`.

```js
assert.deepEqual(cardValues(["8D", "QS", "2H", "JC", "9C"]), [8, 10, 2, 10, 9]);
```

`cardValues(["AS", "KS"])` має повернути `[1, 10]`.

```js
assert.deepEqual(cardValues(["AS", "KS"]), [1, 10]);
```

`cardValues(["10H", "JH", "QH", "KH", "AH"])` має повернути `[10, 10, 10, 10, 1]`.

```js
assert.deepEqual(cardValues(["10H", "JH", "QH", "KH", "AH"]), [10, 10, 10, 10, 1]);
```

# --seed--

## --seed-contents--

```js
function cardValues(cards) {

  return cards;
}
```

# --solutions--

```js
function cardValues(cards) {
  const values = [];

  for (let card of cards) {
    let valueStr = card.slice(0, card.length - 1);

    let value;
    if (valueStr === "A") {
      value = 1;
    } else if (["J", "Q", "K"].includes(valueStr)) {
      value = 10;
    } else {
      value = parseInt(valueStr, 10);
    }

    values.push(value);
  }

  return values;
}
```
