---
id: 6994cff2290543b3aec9f511
title: "Завдання 212: вставлення в масив"
challengeType: 28
dashedName: challenge-212
---

# --description--

Дано масив, значення для вставлення в масив і індекс, за яким потрібно вставити це значення. Поверніть новий масив зі вставленим значенням на вказаному індексі.

# --hints--

`insertIntoArray([2, 4, 8, 10], 6, 2)` має повернути `[2, 4, 6, 8, 10]`.

```js
assert.deepEqual(insertIntoArray([2, 4, 8, 10], 6, 2), [2, 4, 6, 8, 10]);
```

`insertIntoArray(["the", "quick", "fox"], "brown", 2)` має повернути `["the", "quick", "brown", "fox"]`.

```js
assert.deepEqual(insertIntoArray(["the", "quick", "fox"], "brown", 2), ["the", "quick", "brown", "fox"]);
```

`insertIntoArray([], 0, 0)` має повернути `[0]`.

```js
assert.deepEqual(insertIntoArray([], 0, 0), [0]);
```

`insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5)` має повернути `[0, 1, 1, 2, 3, 5, 8, 13]`.

```js
assert.deepEqual(insertIntoArray([0, 1, 1, 2, 3, 8, 13], 5, 5), [0, 1, 1, 2, 3, 5, 8, 13]);
```

# --seed--

## --seed-contents--

```js
function insertIntoArray(arr, value, index) {

  return arr;
}
```

# --solutions--

```js
function insertIntoArray(arr, value, index) {
  return [
    ...arr.slice(0, index),
    value,
    ...arr.slice(index)
  ];
}
```
