---
id: 6994cff2290543b3aec9f50f
title: "Завдання 210: валідатор HSL"
challengeType: 28
dashedName: challenge-210
---

# --description--

Дано рядок, визначте, чи є він дійсним CSS `hsl()` значенням кольору.

- Дійсне значення HSL має бути у форматі `"hsl(h, s%, l%)"`, де:
  - `h` (відтінок) має бути числом від 0 до 360 (включно).
  - `s` (насиченість) має бути відсотком від 0% до 100%.
  - `l` (світлість) має бути відсотком від 0% до 100%.

- Пробіли дозволені лише:
  - Після відкриваючої дужки
  - Перед і/або після ком
  - Перед і/або після закриваючої дужки

- За бажанням, значення може закінчуватися крапкою з комою (`";"`).

Наприклад, `"hsl(240, 50%, 50%)"` є дійсним значенням HSL.

# --hints--

`isValidHSL("hsl(240, 50%, 50%)")` має повертати `true`.

```js
assert.isTrue(isValidHSL("hsl(240, 50%, 50%)"));
```

`isValidHSL("hsl( 200 , 50% , 75% )")` має повертати `true`.

```js
assert.isTrue(isValidHSL("hsl( 200 , 50% , 75% )"));
```

`isValidHSL("hsl(99,60%,80%);")` має повертати `true`.

```js
assert.isTrue(isValidHSL("hsl(99,60%,80%);"));
```

`isValidHSL("hsl(0, 0%, 0%) ;")` має повертати `true`.

```js
assert.isTrue(isValidHSL("hsl(0, 0%, 0%) ;"));
```

`isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;")` має повертати `true`.

```js
assert.isTrue(isValidHSL("hsl(  10  ,  20%   ,  30%   )    ;"));
```

`isValidHSL("hsl(361, 50%, 80%)")` має повертати `false`.

```js
assert.isFalse(isValidHSL("hsl(361, 50%, 80%)"));
```

`isValidHSL("hsl(300, 101%, 70%)")` має повертати `false`.

```js
assert.isFalse(isValidHSL("hsl(300, 101%, 70%)"));
```

`isValidHSL("hsl(200, 55%, 75)")` має повертати `false`.

```js
assert.isFalse(isValidHSL("hsl(200, 55%, 75)"));
```

`isValidHSL("hsl (80, 20%, 10%)")` має повертати `false`.

```js
assert.isFalse(isValidHSL("hsl (80, 20%, 10%)"));
```

# --seed--

## --seed-contents--

```js
function isValidHSL(hsl) {

  return hsl;
}
```

# --solutions--

```js
function isValidHSL(hsl) {
  const regex = /^hsl\(\s*(\d{1,3})\s*,\s*(\d{1,3})%\s*,\s*(\d{1,3})%\s*\)\s*;?$/;

  const match = hsl.match(regex);
  if (!match) return false;

  const hue = Number(match[1]);
  const sat = Number(match[2]);
  const light = Number(match[3]);

  if (hue < 0 || hue > 360) return false;
  if (sat < 0 || sat > 100) return false;
  if (light < 0 || light > 100) return false;

  return true;
}
```
