---
id: 699c8e045ee7cb94ed2322d8
title: "Завдання 217: захоплені шахові фігури"
challengeType: 28
dashedName: challenge-217
---

# --description--

Дано масив рядків, які представляють шахові фігури, що залишилися у вас на дошці. Обчисліть загальне значення фігур, які захопив ваш суперник.

У шахах ви починаєте з 16 фігур:

| Фігура | Скорочення    | Кількість | Значення |
| ------ | ------------- | --------- | -------- |
| Пішак  | `"P"`     | 8         | 1        |
| Тура   | `"R"`     | 2         | 5        |
| Кінь   | `"N"`     | 2         | 3        |
| Слон   | `"B"`     | 2         | 3        |
| Ферзь  | `"Q"`     | 1         | 9        |
| Король | `"K"`     | 1         | 0        |

- Даний масив міститиме лише наведені вище скорочення.
- Будь-які з 16 фігур, яких немає у цьому масиві, були захоплені.
- Поверніть загальне значення всіх захоплених фігур, якщо...
- Якщо захоплено Короля, поверніть `"Checkmate"`.

# --hints--

`getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"])` має повертати `8`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "P", "R", "R", "N", "B", "Q", "K"]), 8);
```

`getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"])` має повертати `26`.

```js
assert.equal(getCapturedValue(["P", "P", "P", "P", "P", "R", "B", "K"]), 26);
```

`getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"])` має повертати `16`.

```js
assert.equal(getCapturedValue(["K", "P", "P", "N", "P", "P", "R", "P", "B", "P", "N", "B"]), 16);
```

`getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"])` має повертати `4`.

```js
assert.equal(getCapturedValue(["P", "Q", "N", "P", "P", "B", "K", "P", "R", "R", "P", "P", "B", "P"]), 4);
```

`getCapturedValue(["P", "K"])` має повертати `38`.

```js
assert.equal(getCapturedValue(["P", "K"]), 38);
```

`getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"])` має повертати `0`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "K", "P", "Q", "N", "P", "P", "R", "R", "P", "P", "P", "B"]), 0);
```

`getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"])` має повертати `"Checkmate"`.

```js
assert.equal(getCapturedValue(["N", "P", "P", "B", "P", "R", "Q", "P", "P", "P", "B"]), "Checkmate");
```

# --seed--

## --seed-contents--

```js
function getCapturedValue(pieces) {

  return pieces;
}
```

# --solutions--

```js
function getCapturedValue(pieces) {
  if (!pieces.includes("K")) return "Checkmate";

  const values = { P:1, R:5, N:3, B:3, Q:9, K:0 };
  let remaining = "PPPPPPPPRRNNBBQK";

  for (const p of pieces) {
    remaining = remaining.replace(p, "");
  }

  return [...remaining].reduce((sum, p) => sum + values[p], 0);
}
```
