---
id: 699c8e045ee7cb94ed2322da
title: "Завдання 219: ювілейні віхи"
challengeType: 28
dashedName: challenge-219
---

# --description--

Дано ціле число, що представляє кількість років, які пара прожила в шлюбі. Поверніть їхній найновіший ювілей відповідно до цієї діаграми:

| Років у шлюбі | Віхи        |
| ------------- | ------------ |
| 1             | `"Paper"`    |
| 5             | `"Wood"`     |
| 10            | `"Tin"`      |
| 25            | `"Silver"`   |
| 40            | `"Ruby"`     |
| 50            | `"Gold"`     |
| 60            | `"Diamond"`  |
| 70            | `"Platinum"` |

- Якщо вони не досягли першої віхи, поверніть `"Newlyweds"`.

# --hints--

`getMilestone(0)` має повертати `"Newlyweds"`.

```js
assert.equal(getMilestone(0), "Newlyweds");
```

`getMilestone(1)` має повертати `"Paper"`.

```js
assert.equal(getMilestone(1), "Paper");
```

`getMilestone(8)` має повертати `"Wood"`.

```js
assert.equal(getMilestone(8), "Wood");
```

`getMilestone(10)` має повертати `"Tin"`.

```js
assert.equal(getMilestone(10), "Tin");
```

`getMilestone(26)` має повертати `"Silver"`.

```js
assert.equal(getMilestone(26), "Silver");
```

`getMilestone(45)` має повертати `"Ruby"`.

```js
assert.equal(getMilestone(45), "Ruby");
```

`getMilestone(50)` має повертати `"Gold"`.

```js
assert.equal(getMilestone(50), "Gold");
```

`getMilestone(64)` має повертати `"Diamond"`.

```js
assert.equal(getMilestone(64), "Diamond");
```

`getMilestone(71)` має повертати `"Platinum"`.

```js
assert.equal(getMilestone(71), "Platinum");
```

# --seed--

## --seed-contents--

```js
function getMilestone(years) {

  return years;
}
```

# --solutions--

```js
function getMilestone(years) {
  if (years < 1) return "Newlyweds";

  const milestones = [
    [1, "Paper"],
    [5, "Wood"],
    [10, "Tin"],
    [25, "Silver"],
    [40, "Ruby"],
    [50, "Gold"],
    [60, "Diamond"],
    [70, "Platinum"]
  ];

  for (let i = milestones.length - 1; i >= 0; i--) {
    if (years >= milestones[i][0]) {
      return milestones[i][1];
    }
  }
}
```
