---
id: 699c8e045ee7cb94ed2322d6
title: "Завдання 215: калькулятор плати за паркування"
challengeType: 28
dashedName: challenge-215
---

# --description--

Дано два рядки, які представляють час, коли ви припаркували автомобіль, і час, коли його забрали. Обчисліть плату за паркування.

- Дано рядки у форматі `"HH:MM"` з використанням 24-годинного формату. Наприклад, `"14:00"` — це 14:00.
- Перший рядок — це час, коли ви припаркували автомобіль, а другий — час, коли його забрали.
- Якщо час забрання раніше за час паркування, це означає, що паркування тривало після опівночі до наступного дня.

Правила плати:

- Кожна година паркування коштує 3 долари.
- Часткові години округлюються вгору до повної години.
- Якщо паркування триває вночі (після опівночі), додається додаткова плата за ніч у розмірі 10 доларів.
- Мінімальна плата становить 5 доларів (використовується лише якщо загальна сума менша за 5).

Поверніть загальну вартість у форматі `"$cost"`, наприклад `"$5"`.

# --hints--

`calculateParkingFee("09:00", "11:00")` має повертати `"$6"`.

```js
assert.equal(calculateParkingFee("09:00", "11:00"), "$6");
```

`calculateParkingFee("10:00", "10:30")` має повертати `"$5"`.

```js
assert.equal(calculateParkingFee("10:00", "10:30"), "$5");
```

`calculateParkingFee("08:10", "10:45")` має повертати `"$9"`.

```js
assert.equal(calculateParkingFee("08:10", "10:45"), "$9");
```

`calculateParkingFee("14:40", "23:10")` має повертати `"$27"`.

```js
assert.equal(calculateParkingFee("14:40", "23:10"), "$27");
```

`calculateParkingFee("18:15", "01:30")` має повертати `"$34"`.

```js
assert.equal(calculateParkingFee("18:15", "01:30"), "$34");
```

`calculateParkingFee("11:11", "11:10")` має повертати `"$82"`.

```js
assert.equal(calculateParkingFee("11:11", "11:10"), "$82");
```

# --seed--

## --seed-contents--

```js
function calculateParkingFee(parkTime, pickupTime) {

  return parkTime;
}
```

# --solutions--

```js
function calculateParkingFee(parkTime, pickupTime) {
  function toMinutes(time) {
    const [hours, minutes] = time.split(":").map(Number);
    return hours * 60 + minutes;
  }

  const entryMinutes = toMinutes(parkTime);
  const exitMinutes = toMinutes(pickupTime);

  let totalMinutes;
  let overnight = false;

  if (exitMinutes < entryMinutes) {
    overnight = true;
    totalMinutes = (24 * 60 - entryMinutes) + exitMinutes;
  } else {
    totalMinutes = exitMinutes - entryMinutes;
  }

  const hours = Math.ceil(totalMinutes / 60);
  let cost = hours * 3;

  if (overnight) {
    cost += 10;
  }

  if (cost < 5) {
    cost = 5;
  }

  return `$${cost}`;
}
```
