---
id: 699c8e045ee7cb94ed2322dd
title: "Завдання 222: тіні в день рівнодення"
challengeType: 28
dashedName: challenge-222
---

# --description--

Сьогодні рівнодення, коли сонце знаходиться безпосередньо над екватором і точно в зеніті опівдні. За заданим часом визначте тінь, яку відкидає вертикальний шестигранний стовп висотою 4 фути.

- Час буде рядком у форматі `"HH:MM"` 24-годинного часу (наприклад, `"15:00"` — це 15:00).
- Час задається лише з кроком у 30 хвилин.

Правила:

- Сонце сходить о 6 ранку безпосередньо `"east"`, а заходить о 6 вечора безпосередньо `"west"`.
- Тінь завжди спрямована в протилежний бік від сонця.
- Довжина тіні (у футах) дорівнює кубу кількості годин, що віддаляють від опівдня.
- Тіні немає до сходу сонця (до 6 ранку), після заходу (6 вечора або пізніше) та опівдні.

Поверніть:

- Якщо тінь існує, поверніть `"(length)ft (direction)"`. Наприклад, `"8ft west"`.
- Інакше поверніть `"No shadow"`.

Наприклад, для `"10:00"` поверніть `"8ft west"`, бо 10 ранку — це 2 години від опівдня, отже 2<sup>3</sup> = 8 футів, а тінь спрямована на захід, бо сонце о 10 ранку знаходиться на сході.

# --hints--

`getShadow("10:00")` має повернути `"8ft west"`.

```js
assert.equal(getShadow("10:00"), "8ft west");
```

`getShadow("15:00")` має повернути `"27ft east"`.

```js
assert.equal(getShadow("15:00"), "27ft east");
```

`getShadow("12:00")` має повернути `"No shadow"`.

```js
assert.equal(getShadow("12:00"), "No shadow");
```

`getShadow("17:30")` має повернути `"166.375ft east"`.

```js
assert.equal(getShadow("17:30"), "166.375ft east");
```

`getShadow("05:00")` має повернути `"No shadow"`.

```js
assert.equal(getShadow("05:00"), "No shadow");
```

`getShadow("06:00")` має повернути `"216ft west"`.

```js
assert.equal(getShadow("06:00"), "216ft west");
```

`getShadow("18:00")` має повернути `"No shadow"`.

```js
assert.equal(getShadow("18:00"), "No shadow");
```

`getShadow("07:30")` має повернути `"91.125ft west"`.

```js
assert.equal(getShadow("07:30"), "91.125ft west");
```

`getShadow("00:00")` має повернути `"No shadow"`.

```js
assert.equal(getShadow("00:00"), "No shadow");
```

# --seed--

## --seed-contents--

```js
function getShadow(time) {

  return time;
}
```

# --solutions--

```js
function getShadow(time) {
  const [hourStr, minuteStr] = time.split(":");
  const hours = Number(hourStr);
  const minutes = Number(minuteStr);
  const timeInHours = hours + minutes / 60;

  if (timeInHours < 6 || timeInHours >= 18 || timeInHours === 12) {
    return "No shadow";
  }

  const hoursFromNoon = Math.abs(12 - timeInHours);
  const length = Math.pow(hoursFromNoon, 3);
  const direction = timeInHours < 12 ? "west" : "east";

  return `${length}ft ${direction}`;
}
```
