---
id: 699c8e045ee7cb94ed2322db
title: "Завдання 220: найбільше число"
challengeType: 28
dashedName: challenge-220
---

# --description--

Дано рядок чисел, розділених різними розділовими знаками. Поверніть найбільше число.

- Даний рядок міститиме лише числа та роздільники.
- Роздільниками можуть бути коми (`","`), знаки оклику (`"!"`), знаки питання (`"?"`), двокрапки (`":"`) або крапки з комою (`";"`).

# --hints--

`largestNumber("1,2")` має повернути `2`.

```js
assert.equal(largestNumber("1,2"), 2);
```

`largestNumber("4;15:60,26?52!0")` має повернути `60`.

```js
assert.equal(largestNumber("4;15:60,26?52!0"), 60);
```

`largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011")` має повернути `-402`.

```js
assert.equal(largestNumber("-402,-1032!-569:-947;-633?-800!-1012;-402,-723?-8102!-3011"), -402);
```

`largestNumber("12;-50,99.9,49.1!-10.1?88?16")` має повернути `99.9`.

```js
assert.equal(largestNumber("12;-50,99.9,49.1!-10.1?88?16"), 99.9);
```

# --seed--

## --seed-contents--

```js
function largestNumber(str) {

  return str;
}
```

# --solutions--

```js
function largestNumber(str) {
  const numbers = str.split(/[,!?:;]/).map(Number);
  return Math.max(...numbers);
}
```
