---
id: 699c8e045ee7cb94ed2322d5
title: "Завдання 214: перевірка ланцюжка доміно"
challengeType: 28
dashedName: challenge-214
---

# --description--

Дано двовимірний масив, який представляє послідовність доміно. Визначте, чи утворює він дійсний ланцюжок.

- Кожен елемент масиву представляє доміно і є масивом із двох чисел від 1 до 6 (включно).
- Щоб ланцюжок був дійсним, друге число кожного доміно має співпадати з першим числом наступного доміно.
- Перше число першого доміно і останнє число останнього доміно не повинні співпадати з будь-яким іншим числом.

# --hints--

`isValidDominoChain([[1, 3], [3, 6], [6, 5]])` має повертати `true`.

```js
assert.isTrue(isValidDominoChain([[1, 3], [3, 6], [6, 5]]));
```

`isValidDominoChain([[6, 2], [3, 4], [4, 1]])` має повертати `false`.

```js
assert.isFalse(isValidDominoChain([[6, 2], [3, 4], [4, 1]]));
```

`isValidDominoChain([[2, 5], [5, 6], [5, 1]])` має повертати `false`.

```js
assert.isFalse(isValidDominoChain([[2, 5], [5, 6], [5, 1]]));
```

`isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]])` має повертати `true`.

```js
assert.isTrue(isValidDominoChain([[4, 3], [3, 1], [1, 6], [6, 6], [6, 5], [5, 1], [1, 1], [1, 4], [4, 4], [4, 2]]));
```

`isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]])` має повертати `false`.

```js
assert.isFalse(isValidDominoChain([[2, 3], [3, 3], [3, 6], [6, 1], [1, 4], [3, 5], [5, 5], [5, 4], [4, 2], [2, 2]]));
```

# --seed--

## --seed-contents--

```js
function isValidDominoChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function isValidDominoChain(dominoes) {
  for (let i = 0; i < dominoes.length - 1; i++) {
    if (dominoes[i][1] !== dominoes[i+1][0]) return false
  }

  return true;
}
```
