---
id: 69a890af247de743333bd4d1
title: "Завдання 228: Кіноніч"
challengeType: 28
dashedName: challenge-228
---

# --description--

Дано рядок із днем тижня, інший рядок із часом сеансу та ціле число квитків. Потрібно повернути загальну вартість квитків на цей сеанс.

Даний день буде одним із:

- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`
- `"Sunday"`

Час сеансу буде заданий у форматі `"H:MMam"` або `"H:MMpm"`. Наприклад `"10:00am"` або `"10:00pm"`.

Поверніть загальну вартість у форматі `"$D.CC"` за такими правилами:

- Вихідні (п’ятниця — неділя): $12.00 за квиток.
- Будні (понеділок — четвер): $10.00 за квиток.
- Денний сеанс (до 17:00): відніміть $2.00 за квиток (крім вівторків).
- Вівторки: усі квитки по $5.00 кожен.

# --hints--

`getMovieNightCost("Saturday", "10:00pm", 1)` має повернути `"$12.00"`.

```js
assert.equal(getMovieNightCost("Saturday", "10:00pm", 1), "$12.00");
```

`getMovieNightCost("Sunday", "10:00am", 1)` має повернути `"$10.00"`.

```js
assert.equal(getMovieNightCost("Sunday", "10:00am", 1), "$10.00");
```

`getMovieNightCost("Tuesday", "7:20pm", 2)` має повернути `"$10.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "7:20pm", 2), "$10.00");
```

`getMovieNightCost("Wednesday", "5:40pm", 3)` має повернути `"$30.00"`.

```js
assert.equal(getMovieNightCost("Wednesday", "5:40pm", 3), "$30.00");
```

`getMovieNightCost("Monday", "11:50am", 4)` має повернути `"$32.00"`.

```js
assert.equal(getMovieNightCost("Monday", "11:50am", 4), "$32.00");
```

`getMovieNightCost("Friday", "4:30pm", 5)` має повернути `"$50.00"`.

```js
assert.equal(getMovieNightCost("Friday", "4:30pm", 5), "$50.00");
```

`getMovieNightCost("Tuesday", "11:30am", 1)` має повернути `"$5.00"`.

```js
assert.equal(getMovieNightCost("Tuesday", "11:30am", 1), "$5.00");
```

# --seed--

## --seed-contents--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {

  return day;
}
```

# --solutions--

```js
function getMovieNightCost(day, showtime, numberOfTickets) {
  let price;

  if (day === "Tuesday") {
    price = 5;
  } else {
    if (["Friday", "Saturday", "Sunday"].includes(day)) {
      price = 12;
    } else {
      price = 10;
    }

    const hour = parseInt(showtime.split(":")[0]);
    const period = showtime.slice(-2);

    let hour24 = hour;

    if (period === "pm" && hour !== 12) hour24 += 12;
    if (period === "am" && hour === 12) hour24 = 0;

    if (hour24 < 17) {
      price -= 2;
    }
  }

  const total = price * numberOfTickets;

  return `$${total.toFixed(2)}`;
}
```
