---
id: 69a890af247de743333bd4d0
title: "Завдання 227: час охолодження"
challengeType: 28
dashedName: challenge-227
---

# --description--

Дано два часові позначки: перша позначає, коли користувач закінчив екзамен, а друга — поточний час. Визначте, чи може користувач знову складати екзамен.

- Обидві часові позначки будуть у форматі `"YYYY-MM-DDTHH:MM:SS"`, наприклад `"2026-03-25T14:00:00"`. Зверніть увагу, що час подано у 24-годинному форматі.
- Користувач повинен чекати щонайменше 48 годин перед повторним складанням екзамену.

# --hints--

`canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00")` має повертати `true`.

```js
assert.isTrue(canRetake("2026-03-23T08:00:00", "2026-03-25T14:00:00"));
```

`canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00")` має повертати `false`.

```js
assert.isFalse(canRetake("2026-03-24T14:00:00", "2026-03-25T10:00:00"));
```

`canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00")` має повертати `true`.

```js
assert.isTrue(canRetake("2026-03-23T09:25:00", "2026-03-25T09:25:00"));
```

`canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59")` має повертати `false`.

```js
assert.isFalse(canRetake("2026-03-23T11:50:00", "2026-03-25T11:49:59"));
```

# --seed--

## --seed-contents--

```js
function canRetake(finishTime, currentTime) {

  return finishTime;
}
```

# --solutions--

```js
function canRetake(finishTime, currentTime) {
  function parse(ts) {
    const parts = ts.split(/[-T:]/).map(Number);
    const [y, m, d, h, min, s] = parts;
    return Date.UTC(y, m - 1, d, h, min, s);
  }

  const finishMs = parse(finishTime);
  const currentMs = parse(currentTime);

  const cooldown = 48 * 60 * 60 * 1000;

  return currentMs - finishMs >= cooldown;
}
```
