---
id: 69a890af247de743333bd4cd
title: "Завдання 224: детектор обсмаження кави"
challengeType: 28
dashedName: challenge-224
---

# --description--

Дано рядок, що представляє зерна, використані для приготування чашки кави. Визначте ступінь обсмаження цієї чашки.

- Даний рядок міститиме такі символи, кожен з яких представляє тип зерна:

  - Апостроф (`'`) — це зерно світлого обсмаження, вартість якого 1 бал за кожне.
  - Дефіс (`-`) — це зерно середнього обсмаження, вартість якого 2 бали за кожне.
  - Крапка (`.`) — це зерно темного обсмаження, вартість якого 3 бали за кожне.

- Рівень обсмаження визначається середнім значенням усіх зерен.

Поверніть:

- `"Light"`, якщо середнє менше за 1.75.
- `"Medium"`, якщо середнє від 1.75 до 2.5.
- `"Dark"`, якщо середнє більше за 2.5.

# --hints--

`detectRoast("''-''''''-'-''--''''")` має повертати `"Light"`.

```js
assert.equal(detectRoast("''-''''''-'-''--''''"), "Light");
```

`detectRoast(".'-''-''..'''.-.-''-")` має повертати `"Medium"`.

```js
assert.equal(detectRoast(".'-''-''..'''.-.-''-"), "Medium");
```

`detectRoast("--.''--'-''.--..-.--")` має повертати `"Medium"`.

```js
assert.equal(detectRoast("--.''--'-''.--..-.--"), "Medium");
```

`detectRoast("-...'-......-..-...-")` має повертати `"Dark"`.

```js
assert.equal(detectRoast("-...'-......-..-...-"), "Dark");
```

`detectRoast(".--.-..-......----.'")` має повертати `"Medium"`.

```js
assert.equal(detectRoast(".--.-..-......----.'"), "Medium");
```

`detectRoast("..-..-..-..-....-.-.")` має повертати `"Dark"`.

```js
assert.equal(detectRoast("..-..-..-..-....-.-."), "Dark");
```

`detectRoast("-'-''''''..-'.''-'.'")` має повертати `"Light"`.

```js
assert.equal(detectRoast("-'-''''''..-'.''-'.'"), "Light");
```

# --seed--

## --seed-contents--

```js
function detectRoast(beans) {

  return beans;
}
```

# --solutions--

```js
function detectRoast(beans) {
  let total = 0;

  for (const bean of beans) {
    if (bean === "'") total += 1;
    else if (bean === "-") total += 2;
    else if (bean === ".") total += 3;
  }

  const avg = total / beans.length;

  if (avg < 1.75) return "Light";
  if (avg <= 2.5) return "Medium";
  return "Dark";
}
```
