---
id: 69a890af247de743333bd4cc
title: "Завдання 223: декодер QR"
challengeType: 28
dashedName: challenge-223
---

# --description--

Дано матрицю 6x6 (масив масивів), що представляє QR-код. Поверніть рядок бінарних даних у коді.

- QR-код може бути повернутий на будь-який кут, кратний 90 градусам.
- Правильно орієнтований код має групу 2x2 `1` (маркерів орієнтації) у нижньому лівому, верхньому лівому та верхньому правому кутах.
- Три маркери орієнтації 2x2 не є частиною бінарних даних.
- Бінарні дані читаються зліва направо, зверху вниз (як книга), коли QR-код правильно орієнтований.
- Код завжди має рівно одну правильну орієнтацію.

Наприклад, для такого вхідного значення:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "110000",
  "110001"
]
```

або для того самого коду з іншою орієнтацією:

```js
[
  "110011",
  "110011",
  "000000",
  "000000",
  "000011",
  "100011"
]
```

Поверніть `"000000000000000000000001"`, усі бінарні дані, крім трьох маркерів орієнтації 2x2.

# --hints--

`decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"])` має повертати `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["110011", "110011", "000000", "000000", "110000", "110001"]), "000000000000000000000001");
```

`decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"])` має повертати `"000000000000000000000001"`.

```js
assert.equal(decodeQr(["100011", "000011", "000000", "000000", "110011", "110011"]), "000000000000000000000001");
```

`decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"])` має повертати `"001101000011000000110100"`.

```js
assert.equal(decodeQr(["110011", "111111", "010000", "110000", "110011", "110100"]), "001101000011000000110100");
```

`decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"])` має повертати `"010001000100010101010110"`.

```js
assert.equal(decodeQr(["011011", "101011", "101000", "100010", "110011", "111011"]), "010001000100010101010110");
```

`decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"])` має повертати `"010000100100100101001110"`.

```js
assert.equal(decodeQr(["111100", "110001", "100011", "001101", "110011", "110011"]), "010000100100100101001110");
```

# --seed--

## --seed-contents--

```js
function decodeQr(qrCode) {

  return qrCode;
}
```

# --solutions--

```js
function isOriented(code) {
  const markerIndices = [
    [0, 0],
    [0, 1],
    [0, 4],
    [0, 5],
    [1, 0],
    [1, 1],
    [1, 4],
    [1, 5],
    [4, 0],
    [4, 1],
    [5, 0],
    [5, 1]
  ];

  for(const [row, col] of markerIndices) {
    if (code[row][col] !== '1') return false;
  }

  return true;
}

function rotate(code) {
  const newCode = ['', '', '', '', '', ''];
  for (let i = 0; i < 6; i++) {
    for (let j = 5; j >= 0; j--) {
      newCode[i] += code[j][i];
    }
  }
  return newCode;
}

function getData(code) {
  let data = '';
  data += code[0].substring(2, 4)
  data += code[1].substring(2, 4)
  data += code[2]
  data += code[3]
  data += code[4].substring(2)
  data += code[5].substring(2)
  return data;
}

function decodeQr(qrCode) {
  let code = [...qrCode];

  while(!isOriented(code)) {
    code = rotate(code);
  }

  return getData(code);
}
```
