---
id: 69b1028d6e265413d0198a29
title: "Завдання 230: рядок трикутника Паскаля"
challengeType: 28
dashedName: challenge-230
---

# --description--

Дано ціле число `n`, поверніть `n`-й рядок трикутника Паскаля у вигляді масиву.

У трикутнику Паскаля кожен рядок починається і закінчується одиницею, а кожне внутрішнє значення є сумою двох значень безпосередньо над ним.

Ось перші 5 рядків трикутника:

```js
    1
   1 1
  1 2 1
 1 3 3 1
1 4 6 4 1
```

# --hints--

`pascalRow(5)` має повертати `[1, 4, 6, 4, 1]`.

```js
assert.deepEqual(pascalRow(5), [1, 4, 6, 4, 1]);
```

`pascalRow(3)` має повертати `[1, 2, 1]`.

```js
assert.deepEqual(pascalRow(3), [1, 2, 1]);
```

`pascalRow(1)` має повертати `[1]`.

```js
assert.deepEqual(pascalRow(1), [1]);
```

`pascalRow(10)` має повертати `[1, 9, 36, 84, 126, 126, 84, 36, 9, 1]`.

```js
assert.deepEqual(pascalRow(10), [1, 9, 36, 84, 126, 126, 84, 36, 9, 1]);
```

`pascalRow(15)` має повертати `[1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]`.

```js
assert.deepEqual(pascalRow(15), [1, 14, 91, 364, 1001, 2002, 3003, 3432, 3003, 2002, 1001, 364, 91, 14, 1]);
```

# --seed--

## --seed-contents--

```js
function pascalRow(n) {

  return n;
}
```

# --solutions--

```js
function pascalRow(n) {
  if (n === 1) return [1];

  let row = [1];

  for (let i = 2; i <= n; i++) {
    const next = [1];

    for (let j = 1; j < row.length; j++) {
      next.push(row[j - 1] + row[j]);
    }

    next.push(1);
    row = next;
  }

  return row;
}
```
