---
id: 69b1028d6e265413d0198a2e
title: "Завдання 235: Капіталізований Фібоначчі"
challengeType: 28
dashedName: challenge-235
---

# --description--

Дано рядок, поверніть новий рядок, у якому кожна літера буде написана з великої літери, якщо її індекс є числом Фібоначчі, а в іншому випадку — у нижньому регістрі.

Послідовність Фібоначчі — це послідовність чисел, де кожне число є сумою двох попередніх. Перші 10 чисел у послідовності: `0`, `1`, `1`, `2`, `3`, `5`, `8`, `13`, `21`, `34`.

- Перший символ має індекс `0`.
- Якщо індекс символів, які не є літерами, є числом Фібоначчі, залиште їх без змін.

# --hints--

`capitalizeFibonacci("hello world")` має повернути `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("hello world"), "HELLo woRld");
```

`capitalizeFibonacci("HELLO WORLD")` має повернути `"HELLo woRld"`.

```js
assert.equal(capitalizeFibonacci("HELLO WORLD"), "HELLo woRld");
```

`capitalizeFibonacci("hello, world!")` має повернути `"HELLo, wOrld!"`.

```js
assert.equal(capitalizeFibonacci("hello, world!"), "HELLo, wOrld!");
```

`capitalizeFibonacci("The quick brown fox jumped over the lazy dog.")` має повернути `"THE qUicK broWn fox jUmped over thE lazy dog."`.

```js
assert.equal(capitalizeFibonacci("The quick brown fox jumped over the lazy dog."), "THE qUicK broWn fox jUmped over thE lazy dog.");
```

`capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl.")` має повернути `"LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl."`.

```js
assert.equal(capitalizeFibonacci("Lorem ipsum dolor sit amet, consectetur adipiscing elit. Proin pulvinar ex nibh, vel ullamcorper ligula egestas quis. Integer tincidunt fringilla accumsan. Integer et metus placerat, gravida felis at, pellentesque nisl."), "LOREm ipSum dOlor sit amet, consecTetur adipiscing elit. proin pulvinar ex nibh, vel ullaMcorper ligula egestas quis. integer tincidunt fringillA accumsan. integer et metus placerat, gravida felis at, pellentesque nisl.");
```

# --seed--

## --seed-contents--

```js
function capitalizeFibonacci(str) {

  return str;
}
```

# --solutions--

```js
function capitalizeFibonacci(str) {
  const fibs = new Set([0, 1]);
  let a = 0, b = 1;
  while (b < str.length) {
    [a, b] = [b, a + b];
    fibs.add(b);
  }

  return str
    .split("")
    .map((char, i) => {
      if (!/[a-zA-Z]/.test(char)) return char;
      return fibs.has(i) ? char.toUpperCase() : char.toLowerCase();
    })
    .join("");
}
```
