---
id: 69b1028d6e265413d0198a2a
title: "Завдання 231: валідатор ISBN-10"
challengeType: 28
dashedName: challenge-231
---

# --description--

Дано рядок, визначте, чи є він дійсним ISBN-10.

ISBN-10 складається з дефісів (`"-"`) та 10 інших символів. Після видалення дефісів (`"-"`):

- Перші 9 символів мають бути цифрами, а
- фінальний символ може бути цифрою або літерою `"X"`, яка представляє число 10.

Щоб перевірити валідність:

- Помножте кожну цифру (або значення) на її позицію (помножте першу цифру на 1, другу на 2 і так далі).
- Додайте всі результати разом.
- Якщо сума ділиться на 11 без остачі, то це дійсний ISBN.

# --hints--

`isValidIsbn10("0-306-40615-2")` має повернути `true`.

```js
assert.isTrue(isValidIsbn10("0-306-40615-2"));
```

`isValidIsbn10("0-306-40615-1")` має повернути `false`.

```js
assert.isFalse(isValidIsbn10("0-306-40615-1"));
```

`isValidIsbn10("0-8044-2957-X")` має повернути `true`.

```js
assert.isTrue(isValidIsbn10("0-8044-2957-X"));
```

`isValidIsbn10("X-306-40615-2")` має повернути `false`.

```js
assert.isFalse(isValidIsbn10("X-306-40615-2"));
```

`isValidIsbn10("0-6822-2589-4")` має повернути `true`.

```js
assert.isTrue(isValidIsbn10("0-6822-2589-4"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn10(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn10(str) {
  const isbn = str.replace(/-/g, "");

  if (isbn.length !== 10) return false;

  let sum = 0;

  for (let i = 0; i < 9; i++) {
    const digit = Number(isbn[i]);
    if (Number.isNaN(digit)) return false;

    sum += digit * (i + 1);
  }

  const last = isbn[9];

  if (last === "X") {
    sum += 10 * 10;
  } else {
    const digit = Number(last);
    if (Number.isNaN(digit)) return false;

    sum += digit * 10;
  }

  return sum % 11 === 0;
}
```
