---
id: 69b1028d6e265413d0198a2f
title: "Завдання 236: історія браузера"
challengeType: 28
dashedName: challenge-236
---

# --description--

Дано масив команд браузера, поверніть масив із двох значень: історію у вигляді масиву URL-адрес і індекс поточної сторінки.

Дійсні команди:

- `"URL"` — де URL — це веб-адреса (наприклад, `"freecodecamp.org"`). Переходить за вказаною URL-адресою, додає її до історії на наступну позицію та видаляє будь-яку історію вперед.
- `"Back"` — переходить на попередню сторінку в історії або залишається на поточній, якщо попередньої немає.
- `"Forward"` — переходить на наступну сторінку в історії або залишається на поточній, якщо наступної немає.

Наприклад, для `["freecodecamp.org", "freecodecamp.org/learn", "Back"]` поверніть `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

# --hints--

`getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"])` має повернути `[["freecodecamp.org", "freecodecamp.org/learn"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["freecodecamp.org", "freecodecamp.org/learn", "Back"]), [["freecodecamp.org", "freecodecamp.org/learn"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"])` має повернути `[["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog"]), [["example.com", "example.com/about", "example.com/contact", "example.com/blog"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact",  "example.com/blog", "Back", "Back", "Forward"])` має повернути `[["example.com", "example.com/contact", "example.com/blog"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "example.com/contact", "example.com/blog", "Back", "Back", "Forward"]), [["example.com", "example.com/contact", "example.com/blog"], 1]);
```

`getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"])` має повернути `[["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "example.com/contact", "example.com/blog", "Back", "Back", "Forward", "freecodecamp.org"]), [["example.com", "example.com/about", "example.com/contact", "freecodecamp.org"], 3]);
```

`getBrowserHistory(["example.com", "example.com/about", "Back", "Back"])` має повернути `[["example.com", "example.com/about"], 0]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Back", "Back"]), [["example.com", "example.com/about"], 0]);
```

`getBrowserHistory(["example.com", "example.com/about", "Forward"])` має повернути `[["example.com", "example.com/about"], 1]`.

```js
assert.deepEqual(getBrowserHistory(["example.com", "example.com/about", "Forward"]), [["example.com", "example.com/about"], 1]);
```

# --seed--

## --seed-contents--

```js
function getBrowserHistory(commands) {

  return commands;
}
```

# --solutions--

```js
function getBrowserHistory(commands) {
  const stack = [];
  let index = -1;

  for (const command of commands) {
    if (command === "Back") {
      if (index > 0) index--;
    } else if (command === "Forward") {
      if (index < stack.length - 1) index++;
    } else {
      stack.splice(index + 1);
      stack.push(command);
      index++;
    }
  }

  return [stack, index];
}
```
