---
id: 69b1028d6e265413d0198a2c
title: "Завдання 233: Будильник"
challengeType: 28
dashedName: challenge-233
---

# --description--

Дано рядок, що представляє час, коли ви встановили будильник, і рядок, що представляє час, коли ви фактично прокинулися. Визначте, чи прокинулися ви раніше, вчасно чи запізно.

- Обидва часи будуть задані у форматі 24-годинного часу `"HH:MM"`.

Поверніть:

- `"early"`, якщо ви прокинулися раніше за час будильника.
- `"on time"`, якщо ви прокинулися в час будильника або протягом 10-хвилинного вікна повторного сигналу після будильника.
- `"late"`, якщо ви прокинулися більш ніж через 10 хвилин після часу будильника.

Обидва часи належать одному дню.

# --hints--

`alarmCheck("07:00", "06:45")` має повертати `"early"`.

```js
assert.equal(alarmCheck("07:00", "06:45"), "early");
```

`alarmCheck("06:30", "06:30")` має повертати `"on time"`.

```js
assert.equal(alarmCheck("06:30", "06:30"), "on time");
```

`alarmCheck("08:10", "08:15")` має повертати `"on time"`.

```js
assert.equal(alarmCheck("08:10", "08:15"), "on time");
```

`alarmCheck("09:30", "09:45")` має повертати `"late"`.

```js
assert.equal(alarmCheck("09:30", "09:45"), "late");
```

`alarmCheck("08:15", "08:25")` має повертати `"on time"`.

```js
assert.equal(alarmCheck("08:15", "08:25"), "on time");
```

`alarmCheck("05:45", "05:56")` має повертати `"late"`.

```js
assert.equal(alarmCheck("05:45", "05:56"), "late");
```

`alarmCheck("04:30", "04:00")` має повертати `"early"`.

```js
assert.equal(alarmCheck("04:30", "04:00"), "early");
```

# --seed--

## --seed-contents--

```js
function alarmCheck(alarmTime, wakeTime) {

  return alarmTime;
}
```

# --solutions--

```js
function alarmCheck(alarmTime, wakeTime) {
  const [alarmH, alarmM] = alarmTime.split(":").map(Number);
  const [wakeH, wakeM] = wakeTime.split(":").map(Number);

  const alarmMinutes = alarmH * 60 + alarmM;
  const wakeMinutes = wakeH * 60 + wakeM;
  const diff = wakeMinutes - alarmMinutes;

  if (diff < 0) return "early";
  if (diff <= 10) return "on time";
  return "late";
}
```
