---
id: 69b58ce40693f140c84c8559
title: "Завдання 240: символи паліндрома"
challengeType: 28
dashedName: challenge-240
---

# --description--

Дано рядок, визначте, чи є він паліндромом, і поверніть середній символ (якщо довжина непарна) або два середні символи (якщо довжина парна).

- Паліндром — це рядок, який однаковий зліва направо і справа наліво.
- Якщо це не паліндром, поверніть `"none"`.

# --hints--

`palindromeLocator("racecar")` має повертати `"e"`.

```js
assert.equal(palindromeLocator("racecar"), "e");
```

`palindromeLocator("level")` має повертати `"v"`.

```js
assert.equal(palindromeLocator("level"), "v");
```

`palindromeLocator("freecodecamp")` має повертати `"none"`.

```js
assert.equal(palindromeLocator("freecodecamp"), "none");
```

`palindromeLocator("noon")` має повертати `"oo"`.

```js
assert.equal(palindromeLocator("noon"), "oo");
```

`palindromeLocator("11100111")` має повертати `"00"`.

```js
assert.equal(palindromeLocator("11100111"), "00");
```

# --seed--

## --seed-contents--

```js
function palindromeLocator(str) {

  return str;
}
```

# --solutions--

```js
function palindromeLocator(str) {
  if (str !== str.split("").reverse().join("")) return "none";

  const mid = Math.floor(str.length / 2);
  return str.length % 2 === 1 ? str[mid] : str[mid - 1] + str[mid];
}
```
