---
id: 69b58ce40693f140c84c855b
title: "Завдання 242: наступне число бінго"
challengeType: 28
dashedName: challenge-242
---

# --description--

За заданим числом бінго поверніть наступне число бінго послідовно.

Число бінго — це одна літера, за якою йде число в межах свого діапазону згідно з цією діаграмою:

| Літера | Діапазон чисел |
| - | - |
| `"B"` | 1-15 |
| `"I"` | 16-30 |
| `"N"` | 31-45 |
| `"G"` | 46-60 |
| `"O"` | 61-75 |

Наприклад, якщо задано `"B10"`, поверніть `"B11"`, наступне число бінго. Якщо задано останнє число бінго, поверніть `"B1"`.

# --hints--

`getNextBingoNumber("B10")` має повернути `"B11"`.

```js
assert.equal(getNextBingoNumber("B10"), "B11");
```

`getNextBingoNumber("N33")` має повернути `"N34"`.

```js
assert.equal(getNextBingoNumber("N33"), "N34");
```

`getNextBingoNumber("I30")` має повернути `"N31"`.

```js
assert.equal(getNextBingoNumber("I30"), "N31");
```

`getNextBingoNumber("G60")` має повернути `"O61"`.

```js
assert.equal(getNextBingoNumber("G60"), "O61");
```

`getNextBingoNumber("O75")` має повернути `"B1"`.

```js
assert.equal(getNextBingoNumber("O75"), "B1");
```

# --seed--

## --seed-contents--

```js
function getNextBingoNumber(n) {

  return n;
}
```

# --solutions--

```js
function getNextBingoNumber(n) {
  const num = parseInt(n.slice(1));
  const next = num === 75 ? 1 : num + 1;

  if (next <= 15) return "B" + next;
  if (next <= 30) return "I" + next;
  if (next <= 45) return "N" + next;
  if (next <= 60) return "G" + next;
  return "O" + next;
}
```
