---
id: 69b559d2903b9e4afe9075f7
title: "Завдання 237: валідація рівняння"
challengeType: 28
dashedName: challenge-237
---

# --description--

Дано рядок, що представляє математичне рівняння. Визначте, чи воно правильне.

- Ліва частина може містити до трьох натуральних чисел і оператори `+`, `-`, `*` та `/`.
- Рівняння буде задане у форматі: `"number operator number = number"` (з двома або трьома числами зліва). Наприклад: `"2 + 2 = 4"` або `"2 + 3 - 1 = 4"`.
- Права частина завжди буде одним цілим числом.

Дотримуйтеся стандартного порядку операцій: множення та ділення виконуються перед додаванням і відніманням, зліва направо.

# --hints--

`isValidEquation("2 + 2 = 4")` має повертати `true`.

```js
assert.isTrue(isValidEquation("2 + 2 = 4"));
```

`isValidEquation("2 + 3 - 1 = 4")` має повертати `true`.

```js
assert.isTrue(isValidEquation("2 + 3 - 1 = 4"));
```

`isValidEquation("8 / 2 = 4")` має повертати `true`.

```js
assert.isTrue(isValidEquation("8 / 2 = 4"));
```

`isValidEquation("10 * 5 = 50")` має повертати `true`.

```js
assert.isTrue(isValidEquation("10 * 5 = 50"));
```

`isValidEquation("2 - 2 = 0")` має повертати `true`.

```js
assert.isTrue(isValidEquation("2 - 2 = 0"));
```

`isValidEquation("2 + 9 / 3 = 5")` має повертати `true`.

```js
assert.isTrue(isValidEquation("2 + 9 / 3 = 5"));
```

`isValidEquation("20 - 2 * 3 = 14")` має повертати `true`.

```js
assert.isTrue(isValidEquation("20 - 2 * 3 = 14"));
```

`isValidEquation("2 + 5 = 6")` має повертати `false`.

```js
assert.isFalse(isValidEquation("2 + 5 = 6"));
```

`isValidEquation("10 - 2 * 3 = 24")` має повертати `false`.

```js
assert.isFalse(isValidEquation("10 - 2 * 3 = 24"));
```

`isValidEquation("3 + 9 / 3 = 4")` має повертати `false`.

```js
assert.isFalse(isValidEquation("3 + 9 / 3 = 4"));
```

# --seed--

## --seed-contents--

```js
function isValidEquation(equation) {

  return equation;
}
```

# --solutions--

```js
function isValidEquation(equation) {
  const [left, right] = equation.split(" = ");
  const tokens = left.split(" ");

  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    if (operator === "*" || operator === "/") {
      const leftOperand = Number(tokens[i - 1]);
      const rightOperand = Number(tokens[i + 1]);
      const computed = operator === "*" ? leftOperand * rightOperand : leftOperand / rightOperand;
      tokens.splice(i - 1, 3, String(computed));
      i -= 2;
    }
  }

  let result = Number(tokens[0]);
  for (let i = 1; i < tokens.length - 1; i += 2) {
    const operator = tokens[i];
    const operand = Number(tokens[i + 1]);
    result = operator === "+" ? result + operand : result - operand;
  }

  return result === Number(right);
}
```
