---
id: 69b559d2903b9e4afe9075f9
title: "Завдання 239: Який сьогодні день?"
challengeType: 28
dashedName: challenge-239
---

# --description--

Дано Unix-мітку часу в мілісекундах, поверніть день тижня.

Дійсні дні для повернення:

- `"Sunday"`
- `"Monday"`
- `"Tuesday"`
- `"Wednesday"`
- `"Thursday"`
- `"Friday"`
- `"Saturday"`

Обов’язково ігноруйте часові пояси.

# --hints--

`getDayOfWeek(1775492249000)` має повертати `"Monday"`.

```js
assert.equal(getDayOfWeek(1775492249000), "Monday");
```

`getDayOfWeek(1766246400000)` має повертати `"Saturday"`.

```js
assert.equal(getDayOfWeek(1766246400000), "Saturday");
```

`getDayOfWeek(33791256000000)` має повертати `"Tuesday"`.

```js
assert.equal(getDayOfWeek(33791256000000), "Tuesday");
```

`getDayOfWeek(1773576000000)` має повертати `"Sunday"`.

```js
assert.equal(getDayOfWeek(1773576000000), "Sunday");
```

`getDayOfWeek(0)` має повертати `"Thursday"`.

```js
assert.equal(getDayOfWeek(0), "Thursday");
```

# --seed--

## --seed-contents--

```js
function getDayOfWeek(timestamp) {

  return timestamp;
}
```

# --solutions--

```js
function getDayOfWeek(timestamp) {
  return new Date(timestamp).toLocaleDateString("en-US", { weekday: "long", timeZone: "UTC" });
}
```
