---
id: 69b559d2903b9e4afe9075f8
title: "Завдання 238: екранування обертання цифр"
challengeType: 28
dashedName: challenge-238
---

# --description--

Дано додатне ціле число. Визначте, чи воно або будь-яке з його обертань ділиться націло на кількість його цифр.

Обертання означає перемістити першу цифру в кінець. Наприклад, після одного обертання 123 стає 231.

- Спочатку перевірте обертання `0` (дане число).
- У заданих числах не буде нулів.
- Поверніть перше число обертання, якщо таке знайдеться, або `"none"`, якщо ні.

# --hints--

`getRotation(123)` має повернути `0`.

```js
assert.equal(getRotation(123), 0);
```

`getRotation(13579)` має повернути `3`.

```js
assert.equal(getRotation(13579), 3);
```

`getRotation(24681)` має повернути `"none"`.

```js
assert.equal(getRotation(24681), "none");
```

`getRotation(84138789345)` має повернути `6`.

```js
assert.equal(getRotation(84138789345), 6);
```

# --seed--

## --seed-contents--

```js
function getRotation(n) {

  return n;
}
```

# --solutions--

```js
function getRotation(n) {
  const str = String(n);
  const digitCount = str.length;
  let current = str;

  for (let i = 0; i < digitCount; i++) {
    if (Number(current) % digitCount === 0) return i;
    current = current.slice(1) + current[0];
  }

  return "none";
}
```
