---
id: 69bc6cb30c1d112a2e110a04
title: "Завдання 247: остання літера"
challengeType: 28
dashedName: challenge-247
---

# --description--

Дано рядок. Поверніть літеру з рядка, яка йде останньою в алфавіті.

- Якщо дві або більше літер однакові за останнє місце в алфавіті, поверніть першу з них.
- Ігноруйте всі символи, що не є літерами.

# --hints--

`getLastLetter("world")` має повернути `"w"`.

```js
assert.equal(getLastLetter("world"), "w");
```

`getLastLetter("Hello World")` має повернути `"W"`.

```js
assert.equal(getLastLetter("Hello World"), "W");
```

`getLastLetter("The quick brown fox jumped over the lazy dog.")` має повернути `"z"`.

```js
assert.equal(getLastLetter("The quick brown fox jumped over the lazy dog."), "z");
```

`getLastLetter("HeLl0")` має повернути `"L"`.

```js
assert.equal(getLastLetter("HeLl0"), "L");
```

`getLastLetter("!#$ er@R asd fT.,> 2t0e9")` має повернути `"T"`.

```js
assert.equal(getLastLetter("!#$ er@R asd fT.,> 2t0e9"), "T");
```

# --seed--

## --seed-contents--

```js
function getLastLetter(str) {

  return str;
}
```

# --solutions--

```js
function getLastLetter(str) {
  const letters = str.split('').filter(char => /[a-zA-Z]/.test(char));
  return letters.reduce((last, char) => char.toLowerCase() > last.toLowerCase() ? char : last);
}
```
