---
id: 69bc6cb30c1d112a2e110a07
title: "Завдання 250: прихований ключ"
challengeType: 28
dashedName: challenge-250
---

# --description--

Ласкаво просимо до 250-го щоденного завдання!

Дано зашифрований рядок, розшифруйте його, використовуючи ключ шифрування, і поверніть результат.

Щоб знайти ключ:

- Подивіться на всі щоденні завдання до сьогоднішнього дня, номер яких кратний 25 (включно з цим).
- Візьміть першу літеру з кожної назви цих завдань і об’єднайте їх у рядок. Якщо назва починається з не літери, знайдіть її першу літеру.

Щоб розшифрувати повідомлення, пройдіться по кожній літері зашифрованого рядка і:

  1. Подивіться на відповідну літеру в ключі (повторюйте ключ, якщо повідомлення довше за ключ).
  2. Перетворіть літеру ключа на відповідне число: `"A"` = 1, `"B"` = 2, ..., `"Z"` = 26.
  3. Зсунути зашифровану літеру назад в алфавіті на це число.
  4. Якщо зсув виходить за `"A"`, оберніться до `"Z"`.

Наприклад, якщо зашифроване повідомлення починається з `"Y"`, а перша літера ключа — `"V"` (22), зсуньте `"Y"` назад на 22 позиції, щоб отримати `"C"`. Повторюйте цей процес для кожної літери, щоб розшифрувати повне повідомлення.

- Зсуваються лише літери, пробіли повертаються без змін.
- Всі задані та повернені літери — у верхньому регістрі.

# --hints--

`decode("YAVJYNXE")` має повернути `"CONGRATS"`.

```js
assert.equal(decode("YAVJYNXE"), "CONGRATS");
```

`decode("YALLUT PQUMJP")` має повернути `"CODING LEGEND"`.

```js
assert.equal(decode("YALLUT PQUMJP"), "CODING LEGEND");
```

`decode("UAC DYR EISAKYM")` має повернути `"YOU ARE AWESOME"`.

```js
assert.equal(decode("UAC DYR EISAKYM"), "YOU ARE AWESOME");
```

`decode("GQMS NBMZU")` має повернути `"KEEP GOING"`.

```js
assert.equal(decode("GQMS NBMZU"), "KEEP GOING");
```

`decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB")` має повернути `"A WINNER IS A DREAMER WHO NEVER GIVES UP"`.

```js
assert.equal(decode("W IQQURV UG I ZDMDTRV IVW JQDHY TMHSA QB"), "A WINNER IS A DREAMER WHO NEVER GIVES UP");
```

# --seed--

## --seed-contents--

```js
function decode(message) {

  return message;
}
```

# --solutions--

```js
function decode(message) {
  const key = "VLHCGMDLNH";
  let keyIndex = 0;
  return message.split('').map(char => {
    if (char === ' ') return char;
    const shift = key[keyIndex % key.length].charCodeAt(0) - 64;
    keyIndex++;
    return String.fromCharCode(((char.charCodeAt(0) - 65 - shift + 26) % 26) + 65);
  }).join('');
}
```
