---
id: 69bc6cb30c1d112a2e110a03
title: "Завдання 246: ініціали імені"
challengeType: 28
dashedName: challenge-246
---

# --description--

Дано повне ім’я у вигляді рядка, поверніть їхні ініціали.

- Імена, з яких потрібно взяти ініціали, розділені пробілом.
- Ініціали мають бути у верхньому регістрі.
- Ініціали мають бути розділені крапками.

Наприклад, `"Tommy Millwood"` повертає `"T.M."`.

# --hints--

`getInitials("Tommy Millwood")` має повернути `"T.M."`.

```js
assert.equal(getInitials("Tommy Millwood"), "T.M.");
```

`getInitials("Savanna Puddlesplash")` має повернути `"S.P."`.

```js
assert.equal(getInitials("Savanna Puddlesplash"), "S.P.");
```

`getInitials("Frances Cowell Conrad")` має повернути `"F.C.C."`.

```js
assert.equal(getInitials("Frances Cowell Conrad"), "F.C.C.");
```

`getInitials("Dragon")` має повернути `"D."`.

```js
assert.equal(getInitials("Dragon"), "D.");
```

`getInitials("Dorothy Vera Clump Haverstock Norris")` має повернути `"D.V.C.H.N."`.

```js
assert.equal(getInitials("Dorothy Vera Clump Haverstock Norris"), "D.V.C.H.N.");
```

# --seed--

## --seed-contents--

```js
function getInitials(name) {

  return name;
}
```

# --solutions--

```js
function getInitials(name) {
  return name.split(' ').map(word => word[0].toUpperCase() + '.').join('');
}
```
