---
id: 69bc6cb30c1d112a2e110a06
title: "Завдання 249: рядкова математика"
challengeType: 28
dashedName: challenge-249
---

# --description--

Дано рядок із числами та іншими символами. Виконайте математичні операції над числами залежно від кількості нецифрових символів між числами.

- Якщо кількість символів між двома числами парна, використовуйте додавання.
- Якщо непарна — використовуйте віднімання.
- Послідовні цифри утворюють одне число.
- Операції виконуються зліва направо.
- Ігноруйте початкові та кінцеві символи, які не є цифрами.

Наприклад, для `"3ab10c8"` поверніть `5`. Додайте 3 і 10, щоб отримати 13, бо між ними парна кількість символів. Потім відніміть 8 від 13, бо між результатом і 8 непарна кількість символів.

# --hints--

`doMath("3ab10c8")` має повернути `5`.

```js
assert.equal(doMath("3ab10c8"), 5);
```

`doMath("6MINUS4")` має повернути `2`.

```js
assert.equal(doMath("6MINUS4"), 2);
```

`doMath("9plus3")` має повернути `12`.

```js
assert.equal(doMath("9plus3"), 12);
```

`doMath("5fkwo#10i#%.<>15P=@20!#B/25")` має повернути `15`.

```js
assert.equal(doMath("5fkwo#10i#%.<>15P=@20!#B/25"), 15);
```

`doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt")` має повернути `67`.

```js
assert.equal(doMath("a.67,1$lk6ldf34@#LD@]2d32d2'2l3,@l3L#@2gh35s09if=df#$t9sm49t0df3$^%[vc;:0:4mt"), 67);
```

# --seed--

## --seed-contents--

```js
function doMath(str) {

  return str;
}
```

# --solutions--

```js
function doMath(str) {
  const tokens = str.match(/(\d+)|(\D+)/g);
  let result = null;
  let pendingOp = null;

  for (const token of tokens) {
    if (/^\d+$/.test(token)) {
      if (result === null) {
        result = parseInt(token);
      } else {
        result = pendingOp === 'add' ? result + parseInt(token) : result - parseInt(token);
      }
    } else {
      pendingOp = token.length % 2 === 0 ? 'add' : 'subtract';
    }
  }

  return result;
}
```
