---
id: 69c5f3d787b1725d5f00c8b8
title: "Завдання 254: слова з непарною кількістю літер"
challengeType: 28
dashedName: challenge-254
---

# --description--

Дано рядок слів, поверніть лише ті слова, у яких непарна кількість літер.

- Слова у заданому рядку будуть розділені одним пробілом.
- Поверніть слова, розділені одним пробілом.

# --hints--

`getOddWords("This is a super good test")` має повернути `"a super"`.

```js
assert.equal(getOddWords("This is a super good test"), "a super");
```

`getOddWords("one two three four")` має повернути `"one two three"`.

```js
assert.equal(getOddWords("one two three four"), "one two three");
```

`getOddWords("banana split sundae with rainbow sprinkles on top")` має повернути `"split rainbow sprinkles top"`.

```js
assert.equal(getOddWords("banana split sundae with rainbow sprinkles on top"), "split rainbow sprinkles top");
```

`getOddWords("The quick brown fox jumped over the lazy river")` має повернути `"The quick brown fox the river"`.

```js
assert.equal(getOddWords("The quick brown fox jumped over the lazy river"), "The quick brown fox the river");
```

# --seed--

## --seed-contents--

```js
function getOddWords(str) {

  return str;
}
```

# --solutions--

```js
function getOddWords(str) {
  return str
    .split(" ")
    .filter(word => word.length % 2 === 1)
    .join(" ");
}
```
