---
id: 69c5f3d787b1725d5f00c8ba
title: "Завдання 256: найближчий напрямок часу"
challengeType: 28
dashedName: challenge-256
---

# --description--

Дано два часи, визначте, чи можна дістатися від першого до другого швидше, рухаючись вперед чи назад.

- Часи задані у 24-годинному форматі (`"HH:MM"`)
- Годинник циклічний (23:59 переходить у 00:00 при русі вперед, а 00:00 переходить у 23:59 при русі назад)

Поверніть:

- `"forward"`, якщо рух вперед коротший
- `"backward"`, якщо рух назад коротший
- `"equal"`, якщо обидва напрямки займають однаковий час

# --hints--

`getDirection("10:00", "12:00")` має повертати `"forward"`.

```js
assert.equal(getDirection("10:00", "12:00"), "forward");
```

`getDirection("11:00", "05:00")` має повертати `"backward"`.

```js
assert.equal(getDirection("11:00", "05:00"), "backward");
```

`getDirection("00:00", "12:00")` має повертати `"equal"`.

```js
assert.equal(getDirection("00:00", "12:00"), "equal");
```

`getDirection("15:45", "01:10")` має повертати `"forward"`.

```js
assert.equal(getDirection("15:45", "01:10"), "forward");
```

`getDirection("03:30", "19:50")` має повертати `"backward"`.

```js
assert.equal(getDirection("03:30", "19:50"), "backward");
```

`getDirection("06:30", "18:30")` має повертати `"equal"`.

```js
assert.equal(getDirection("06:30", "18:30"), "equal");
```

# --seed--

## --seed-contents--

```js
function getDirection(time1, time2) {

  return time1;
}
```

# --solutions--

```js
function getDirection(time1, time2) {
  function toMinutes(t) {
    const [h, m] = t.split(":").map(Number);
    return h * 60 + m;
  }

  const start = toMinutes(time1);
  const end = toMinutes(time2);

  const forward = (end - start + 1440) % 1440;
  const backward = (start - end + 1440) % 1440;

  if (forward < backward) return "forward";
  if (backward < forward) return "backward";
  return "equal";
}
```
