---
id: 69c5f3d787b1725d5f00c8bd
title: "Завдання 259: вибух FizzBuzz"
challengeType: 28
dashedName: challenge-259
---

# --description--

Дано ціле число, поверніть кількість кроків, необхідних, щоб перетворити слово `"fizzbuzz"` у рядок з принаймні заданою кількістю `"z"` за такими правилами:

- Починайте з рядка `"fizzbuzz"`.
- Кожного кроку застосовуйте стандартні правила FizzBuzz, використовуючи позицію літери в рядку (перша `"f"` має позицію 1).
  - Якщо позиція літери ділиться на 3, замініть літеру на `"fizz"`
  - Якщо вона ділиться на 5, замініть літеру на `"buzz"`
  - Якщо вона ділиться і на 3, і на 5, замініть літеру на `"fizzbuzz"`

Отже, після 1 кроку `"fizzbuzz"` перетворюється на `"fifizzzbuzzfizzzz"`, який містить 9 `"z"`.

# --hints--

`explodeFizzbuzz(9)` має повернути `1`.

```js
assert.equal(explodeFizzbuzz(9), 1);
```

`explodeFizzbuzz(15)` має повернути `2`.

```js
assert.equal(explodeFizzbuzz(15), 2);
```

`explodeFizzbuzz(51)` має повернути `3`.

```js
assert.equal(explodeFizzbuzz(51), 3);
```

`explodeFizzbuzz(52)` має повернути `4`.

```js
assert.equal(explodeFizzbuzz(52), 4);
```

`explodeFizzbuzz(359)` має повернути `5`.

```js
assert.equal(explodeFizzbuzz(359), 5);
```

`explodeFizzbuzz(789)` має повернути `6`.

```js
assert.equal(explodeFizzbuzz(789), 6);
```

`explodeFizzbuzz(54482)` має повернути `11`.

```js
assert.equal(explodeFizzbuzz(54482), 11);
```

`explodeFizzbuzz(1000000)` має повернути `14`.

```js
assert.equal(explodeFizzbuzz(1000000), 14);
```

# --seed--

## --seed-contents--

```js
function explodeFizzbuzz(targetZCount) {

  return targetZCount;
}
```

# --solutions--

```js
function explodeFizzbuzz(targetZCount) {
  let str = "fizzbuzz";
  let steps = 0;

  while ((str.match(/z/g) || []).length < targetZCount) {
    let next = "";
    for (let i = 0; i < str.length; i++) {
      const pos = i + 1;
      if (pos % 15 === 0) next += "fizzbuzz";
      else if (pos % 3 === 0) next += "fizz";
      else if (pos % 5 === 0) next += "buzz";
      else next += str[i];
    }
    str = next;
    steps++;
  }

  return steps;
}
```
