---
id: 69c5f3d787b1725d5f00c8b7
title: "Завдання 253: пошук абревіатури"
challengeType: 28
dashedName: challenge-253
---

# --description--

Дано рядок, що представляє абревіатуру. Поверніть повну назву організації, якій вона належить, зі списку нижче:

- `"National Avocado Storage Authority"`
- `"Cats Infiltration Agency"`
- `"Fluffy Beanbag Inspectors"`
- `"Department Of Jelly"`
- `"Wild Honey Organization"`
- `"Eating Pancakes Administration"`

Кожна літера в заданій абревіатурі має відповідати першій літері кожного слова в назві організації, якій вона належить, у тому ж порядку.

# --hints--

`findOrg("NASA")` має повернути `"National Avocado Storage Authority"`.

```js
assert.equal(findOrg("NASA"), "National Avocado Storage Authority");
```

`findOrg("CIA")` має повернути `"Cats Infiltration Agency"`.

```js
assert.equal(findOrg("CIA"), "Cats Infiltration Agency");
```

`findOrg("FBI")` має повернути `"Fluffy Beanbag Inspectors"`.

```js
assert.equal(findOrg("FBI"), "Fluffy Beanbag Inspectors");
```

`findOrg("DOJ")` має повернути `"Department Of Jelly"`.

```js
assert.equal(findOrg("DOJ"), "Department Of Jelly");
```

`findOrg("WHO")` має повернути `"Wild Honey Organization"`.

```js
assert.equal(findOrg("WHO"), "Wild Honey Organization");
```

`findOrg("EPA")` має повернути `"Eating Pancakes Administration"`.

```js
assert.equal(findOrg("EPA"), "Eating Pancakes Administration");
```

# --seed--

## --seed-contents--

```js
function findOrg(acronym) {

  return acronym;
}
```

# --solutions--

```js
function findOrg(acronym) {
  const orgs = [
    "National Avocado Storage Authority",
    "Cats Infiltration Agency",
    "Fluffy Beanbag Inspectors",
    "Department Of Jelly",
    "Wild Honey Organization",
    "Eating Pancakes Administration"
  ];

  for (let org of orgs) {
    const initials = org
      .split(" ")
      .map(word => word[0])
      .join("");

    if (initials.toUpperCase() === acronym.toUpperCase()) {
      return org;
    }
  }

  return "not found";
}
```
