---
id: 69c6ff713a52713463aa7926
title: "Завдання 260: оцінка слова"
challengeType: 28
dashedName: challenge-260
---

# --description--

Дано слово, поверніть його оцінку, використовуючи стандартну таблицю значень літер:

| Літера | Значення |
| --- | --- |
|   A |   1 |
|   B |   2 |
| ... | ... |
|   Z |  26 |

- Літери у верхньому та нижньому регістрах мають однакове значення.

# --hints--

`getWordScore("hi")` має повертати `17`.

```js
assert.equal(getWordScore("hi"), 17);
```

`getWordScore("hello")` має повертати `52`.

```js
assert.equal(getWordScore("hello"), 52);
```

`getWordScore("hippopotamus")` має повертати `169`.

```js
assert.equal(getWordScore("hippopotamus"), 169);
```

`getWordScore("freeCodeCamp")` має повертати `94`.

```js
assert.equal(getWordScore("freeCodeCamp"), 94);
```

# --seed--

## --seed-contents--

```js
function getWordScore(word) {

  return word;
}
```

# --solutions--

```js
function getWordScore(word) {
  return word.toLowerCase().split('').reduce((sum, char) => {
    return sum + (char.charCodeAt(0) - 96);
  }, 0);
}
```
