---
id: 69c6ff713a52713463aa7927
title: "Завдання 261: числові слова"
challengeType: 28
dashedName: challenge-261
---

# --description--

Дано ціле число від 0 до 99, поверніть його англійське словесне представлення.

- `0` повертає `"zero"`.
- Числа від 1 до 19 мають унікальні назви (`"one"`, `"two"`, ..., `"ten"`, `"eleven"`, ..., `"eighteen"`, `"nineteen"`).
- Кратні 10 від 20 до 90 мають власні назви (`"twenty"`, `"thirty"`, ..., `"eighty"`, `"ninety"`).
- Числа від 21 до 99, які не є кратними 10, записуються як два слова, з’єднані дефісом. Наприклад `"forty-two"` і `"fifty-three"`.

# --hints--

`getNumberWords(0)` має повертати `"zero"`.

```js
assert.equal(getNumberWords(0), "zero");
```

`getNumberWords(10)` має повертати `"ten"`.

```js
assert.equal(getNumberWords(10), "ten");
```

`getNumberWords(19)` має повертати `"nineteen"`.

```js
assert.equal(getNumberWords(19), "nineteen");
```

`getNumberWords(30)` має повертати `"thirty"`.

```js
assert.equal(getNumberWords(30), "thirty");
```

`getNumberWords(53)` має повертати `"fifty-three"`.

```js
assert.equal(getNumberWords(53), "fifty-three");
```

`getNumberWords(7)` має повертати `"seven"`.

```js
assert.equal(getNumberWords(7), "seven");
```

`getNumberWords(12)` має повертати `"twelve"`.

```js
assert.equal(getNumberWords(12), "twelve");
```

`getNumberWords(60)` має повертати `"sixty"`.

```js
assert.equal(getNumberWords(60), "sixty");
```

`getNumberWords(67)` має повертати `"sixty-seven"`.

```js
assert.equal(getNumberWords(67), "sixty-seven");
```

`getNumberWords(98)` має повертати `"ninety-eight"`.

```js
assert.equal(getNumberWords(98), "ninety-eight");
```

# --seed--

## --seed-contents--

```js
function getNumberWords(n) {

  return n;
}
```

# --solutions--

```js
function getNumberWords(n) {
  const ones = ["zero","one","two","three","four","five","six","seven","eight","nine",
                 "ten","eleven","twelve","thirteen","fourteen","fifteen","sixteen",
                 "seventeen","eighteen","nineteen"];
  const tens = ["","","twenty","thirty","forty","fifty","sixty","seventy","eighty","ninety"];

  if (n < 20) return ones[n];
  if (n % 10 === 0) return tens[Math.floor(n / 10)];
  return tens[Math.floor(n / 10)] + "-" + ones[n % 10];
}
```
