---
id: 69c6ff713a52713463aa7929
title: "Завдання 263: бінарний кросворд"
challengeType: 28
dashedName: challenge-263
---

# --description--

За заданим символом визначте, чи можна знайти його 8-бітне бінарне представлення у наведеній сітці, горизонтально або вертикально в будь-якому напрямку:

```js
0 1 0 0 0 0 0 1
0 1 1 0 1 1 1 1
0 1 0 0 0 1 0 0
0 1 1 0 0 1 0 1
0 1 0 1 0 0 1 0
0 1 0 1 0 1 0 0
0 1 1 0 1 0 0 0
1 0 1 0 1 1 1 0
```

Наприклад, `"A"` має бінарне представлення `01000001`, яке з’являється в першому рядку зліва направо.

# --hints--

`isInCrossword("I")` має повернути `true`.

```js
assert.isTrue(isInCrossword("I"));
```

`isInCrossword("D")` має повернути `true`.

```js
assert.isTrue(isInCrossword("D"));
```

`isInCrossword("0")` має повернути `true`.

```js
assert.isTrue(isInCrossword("0"));
```

`isInCrossword("u")` має повернути `true`.

```js
assert.isTrue(isInCrossword("u"));
```

`isInCrossword("Y")` має повернути `false`.

```js
assert.isFalse(isInCrossword("Y"));
```

`isInCrossword("p")` має повернути `false`.

```js
assert.isFalse(isInCrossword("p"));
```

`isInCrossword("1")` має повернути `false`.

```js
assert.isFalse(isInCrossword("1"));
```

`isInCrossword("Q")` має повернути `false`.

```js
assert.isFalse(isInCrossword("Q"));
```

# --seed--

## --seed-contents--

```js
function isInCrossword(char) {

  return char;
}
```

# --solutions--

```js
function isInCrossword(char) {
  const grid = [
    "01000001",
    "01101111",
    "01000100",
    "01100101",
    "01010010",
    "01010100",
    "01101000",
    "10101110",
  ];

  const target = char.charCodeAt(0).toString(2).padStart(8, '0');
  const cols = Array.from({length: 8}, (_, c) => grid.map(r => r[c]).join(''));

  return [...grid, ...cols].some(seq => seq.includes(target) || seq.split('').reverse().join('').includes(target));
}
```
