---
id: 69c6ff713a52713463aa7928
title: "Завдання 262: парсер URL-запиту"
challengeType: 28
dashedName: challenge-262
---

# --description--

Дано URL, який містить рядок запиту. Проаналізуйте цей рядок запиту та перетворіть його на об’єкт (або словник) пар ключ-значення.

- Рядок запиту починається після `"?"`,
- кожен параметр відокремлений `"&"`,
- кожна пара ключ/значення відокремлена `"="`

Наприклад, для `"https://example.com/search?name=Alice&age=30"` поверніть:

```json
{
  "name": "Alice",
  "age": "30"
}
```

Усі значення мають бути повернені у вигляді рядків.

# --hints--

`parseUrlQuery("https://example.com/search?name=Alice&age=30")` має повернути `{"name": "Alice", "age": "30"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com/search?name=Alice&age=30"), {"name": "Alice", "age": "30"});
```

`parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python")` має повернути `{"skill": "programming", "language": "python"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/learn?skill=programming&language=python"), {"skill": "programming", "language": "python"});
```

`parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2")` має повернути `{"category": "books", "sort": "asc", "page": "2"}`

```js
assert.deepEqual(parseUrlQuery("https://freecodecamp.org/items?category=books&sort=asc&page=2"), {"category": "books", "sort": "asc", "page": "2"});
```

`parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now")` має повернути `{"redirect": "freecodecamp.org/learn", "when": "now"}`

```js
assert.deepEqual(parseUrlQuery("https://example.com?redirect=freecodecamp.org/learn&when=now"), {"redirect": "freecodecamp.org/learn", "when": "now"});
```

# --seed--

## --seed-contents--

```js
function parseUrlQuery(url) {

  return url;
}
```

# --solutions--

```js
function parseUrlQuery(url) {
  const queryString = url.split('?')[1];
  return Object.fromEntries(
    queryString.split('&').map(param => param.split('='))
  );
}
```
