---
id: 69cfca90e8a0a6d4d6871c4f
title: "Завдання 265: найглибші дужки"
challengeType: 28
dashedName: challenge-265
---

# --description--

Дано рядок, що містить збалансовані дужки. Поверніть вміст найглибше вкладених дужок.

- Дужки можуть бути будь-якого з трьох типів: `()`, `[]` та `{}`.
- У вхідних даних завжди буде одна найглибша група.

Наприклад, для `"(hello (world))"` поверніть `"world"`.

# --hints--

`getDeepestBrackets("(hello (world))")` має повернути `"world"`.

```js
assert.equal(getDeepestBrackets("(hello (world))"), "world");
```

`getDeepestBrackets("[outer [inner] outer]")` має повернути `"inner"`.

```js
assert.equal(getDeepestBrackets("[outer [inner] outer]"), "inner");
```

`getDeepestBrackets("{a{b}c{d{e}f}g}")` має повернути `"e"`.

```js
assert.equal(getDeepestBrackets("{a{b}c{d{e}f}g}"), "e");
```

`getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]")` має повернути `"fox"`.

```js
assert.equal(getDeepestBrackets("[the {quick (brown [fox] jumped) over (the) lazy} dog]"), "fox");
```

`getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p")` має повернути `"C"`.

```js
assert.equal(getDeepestBrackets("f[(r)e{e}C{o[(d){e(C)}a]m}]p"), "C");
```

# --seed--

## --seed-contents--

```js
function getDeepestBrackets(str) {

  return str;
}
```

# --solutions--

```js
function getDeepestBrackets(str) {
  let maxDepth = 0, depth = 0, content = '';
  let start = 0;

  for (let i = 0; i < str.length; i++) {
    if ('([{'.includes(str[i])) {
      depth++;
      if (depth > maxDepth) {
        maxDepth = depth;
        start = i + 1;
      }
    } else if (')]}'.includes(str[i])) {
      if (depth === maxDepth) {
        content = str.slice(start, i);
      }
      depth--;
    }
  }
  return content;
}
```
