---
id: 69cfca90e8a0a6d4d6871c52
title: "Завдання 268: нарцисичне число"
challengeType: 28
dashedName: challenge-268
---

# --description--

Дано додатне ціле число, визначте, чи є воно нарцисичним числом.

- Число є нарцисичним, якщо сума кожної його цифри, піднесеної до степеня загальної кількості цифр, дорівнює самому числу.

Наприклад, 153 має 3 цифри, і 1<sup>3</sup> + 5<sup>3</sup> + 3<sup>3</sup> = 153, отже воно є нарцисичним.

# --hints--

`isNarcissistic(153)` має повертати `true`.

```js
assert.isTrue(isNarcissistic(153));
```

`isNarcissistic(154)` має повертати `false`.

```js
assert.isFalse(isNarcissistic(154));
```

`isNarcissistic(371)` має повертати `true`.

```js
assert.isTrue(isNarcissistic(371));
```

`isNarcissistic(512)` має повертати `false`.

```js
assert.isFalse(isNarcissistic(512));
```

`isNarcissistic(9)` має повертати `true`.

```js
assert.isTrue(isNarcissistic(9));
```

`isNarcissistic(11)` має повертати `false`.

```js
assert.isFalse(isNarcissistic(11));
```

`isNarcissistic(9474)` має повертати `true`.

```js
assert.isTrue(isNarcissistic(9474));
```

`isNarcissistic(6549)` має повертати `false`.

```js
assert.isFalse(isNarcissistic(6549));
```

# --seed--

## --seed-contents--

```js
function isNarcissistic(n) {

  return n;
}
```

# --solutions--

```js
function isNarcissistic(n) {
  const digits = String(n).split('');
  const power = digits.length;
  const sum = digits.reduce((acc, d) => acc + Math.pow(Number(d), power), 0);
  return sum === n;
}
```
