---
id: 69cfca90e8a0a6d4d6871c50
title: "Завдання 266: Добрий день"
challengeType: 28
dashedName: challenge-266
---

# --description--

Дано рядок часу у форматі `"HH:MM"` (24-годинний формат), поверніть:

- `"Good morning"` для часу від `05:00` до `11:59`
- `"Good afternoon"` для часу від `12:00` до `17:59`
- `"Good evening"` для часу від `18:00` до `21:59`
- `"Good night"` для часу від `22:00` до `04:59`

# --hints--

`getGreeting("06:30")` має повертати `"Good morning"`.

```js
assert.equal(getGreeting("06:30"), "Good morning");
```

`getGreeting("12:00")` має повертати `"Good afternoon"`.

```js
assert.equal(getGreeting("12:00"), "Good afternoon");
```

`getGreeting("21:59")` має повертати `"Good evening"`.

```js
assert.equal(getGreeting("21:59"), "Good evening");
```

`getGreeting("00:01")` має повертати `"Good night"`.

```js
assert.equal(getGreeting("00:01"), "Good night");
```

`getGreeting("11:30")` має повертати `"Good morning"`.

```js
assert.equal(getGreeting("11:30"), "Good morning");
```

# --seed--

## --seed-contents--

```js
function getGreeting(time) {

  return time;
}
```

# --solutions--

```js
function getGreeting(time) {
  const [hours, minutes] = time.split(':').map(Number);
  const total = hours * 60 + minutes;

  if (total >= 300 && total < 720) return 'Good morning';
  if (total >= 720 && total < 1080) return 'Good afternoon';
  if (total >= 1080 && total < 1320) return 'Good evening';
  return 'Good night';
}
```
