---
id: 69cfca90e8a0a6d4d6871c51
title: "Завдання 267: конвертер парсеків"
challengeType: 28
dashedName: challenge-267
---

# --description--

У далекій галактиці парсеки використовують для вимірювання як часу, так і відстані. За заданим цілим числом парсеків поверніть його еквівалент у часі або відстані.

- Якщо задане ціле число непарне, воно представляє час. Якщо парне — відстань.

Використовуйте такі коефіцієнти конвертації:

| Парсеки | Час/Відстань |
| - | - |
| 1 | 2 години |
| 2 | 6 світлових років |

Поверніть конвертоване значення як ціле число.

# --hints--

`convertParsecs(1)` має повертати `2`.

```js
assert.equal(convertParsecs(1), 2);
```

`convertParsecs(2)` має повертати `6`.

```js
assert.equal(convertParsecs(2), 6);
```

`convertParsecs(31)` має повертати `62`.

```js
assert.equal(convertParsecs(31), 62);
```

`convertParsecs(88)` має повертати `264`.

```js
assert.equal(convertParsecs(88), 264);
```

`convertParsecs(17)` має повертати `34`.

```js
assert.equal(convertParsecs(17), 34);
```

`convertParsecs(14)` має повертати `42`.

```js
assert.equal(convertParsecs(14), 42);
```

# --seed--

## --seed-contents--

```js
function convertParsecs(parsecs) {

  return parsecs;
}
```

# --solutions--

```js
function convertParsecs(parsecs) {
  if (parsecs % 2 !== 0) {
    return parsecs * 2;
  } else {
    return parsecs * 3;
  }
}
```
