---
id: 69e2383af7832c8032603b8f
title: "Завдання 273: валідатор ISBN-13"
challengeType: 28
dashedName: challenge-273
---

# --description--

Дано рядок, визначте, чи є він дійсним номером ISBN-13.

Дійсний ISBN-13:

- Містить лише цифри та дефіси
- Містить рівно 13 цифр після видалення дефісів
- Проходить таку перевірку:
  1. Помножте кожну цифру на 1 або 3, чергуючи (помножте першу цифру на 1, другу на 3, третю на 1 і так далі).
  2. Сума результатів має бути кратною 10.

# --hints--

`isValidIsbn13("9780306406157")` має повертати `true`.

```js
assert.isTrue(isValidIsbn13("9780306406157"));
```

`isValidIsbn13("97803064061570")` має повертати `false`.

```js
assert.isFalse(isValidIsbn13("97803064061570"));
```

`isValidIsbn13("978-0-13-595705-9")` має повертати `true`.

```js
assert.isTrue(isValidIsbn13("978-0-13-595705-9"));
```

`isValidIsbn13("978-030-64061A-4")` має повертати `false`.

```js
assert.isFalse(isValidIsbn13("978-030-64061A-4"));
```

`isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9")` має повертати `true`.

```js
assert.isTrue(isValidIsbn13("9-7-8-0-1-3-4-7-5-7-5-9-9"));
```

# --seed--

## --seed-contents--

```js
function isValidIsbn13(str) {

  return str;
}
```

# --solutions--

```js
function isValidIsbn13(str) {
  const digits = str.replace(/-/g, '');
  if (!/^\d{13}$/.test(digits)) return false;
  const sum = digits.split('').reduce((acc, d, i) => acc + d * (i % 2 === 0 ? 1 : 3), 0);
  return sum % 10 === 0;
}
```
