---
id: 69e2383af7832c8032603b92
title: "Завдання 276: елемент, що порушує порядок"
challengeType: 28
dashedName: challenge-276
---

# --description--

Дано масив цілих чисел, відсортований за зростанням, крім одного елемента, який стоїть не на своєму місці. Потрібно повернути індекс цього елемента.

- Якщо таких елементів кілька, поверніть індекс першого з них.

# --hints--

`findOffender([1, 6, 2, 3, 4, 5])` має повертати `1`.

```js
assert.equal(findOffender([1, 6, 2, 3, 4, 5]), 1);
```

`findOffender([1, 2, 3, 5, 4, 5])` має повертати `3`.

```js
assert.equal(findOffender([1, 2, 3, 5, 4, 5]), 3);
```

`findOffender([2, 1])` має повертати `0`.

```js
assert.equal(findOffender([2, 1]), 0);
```

`findOffender([2, 4, 1, 6, 8])` має повертати `2`.

```js
assert.equal(findOffender([2, 4, 1, 6, 8]), 2);
```

`findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91])` має повертати `6`.

```js
assert.equal(findOffender([5, 18, 24, 33, 40, 55, 15, 68, 84, 91]), 6);
```

# --seed--

## --seed-contents--

```js
function findOffender(arr) {

  return arr;
}
```

# --solutions--

```js
function findOffender(arr) {
  for (let i = 0; i < arr.length; i++) {
    const without = [...arr.slice(0, i), ...arr.slice(i + 1)];
    if (without.every((n, j) => j === 0 || without[j - 1] <= n)) return i;
  }
}
```
