---
id: 69f35a5bb823ed620fcb7cbb
title: "Завдання 282: Сонний борг"
challengeType: 28
dashedName: challenge-282
---

# --description--

Дано масив годин сну щоночі до сьогоднішнього дня та цільове число годин сну за ніч. Поверніть, скільки годин сну вам потрібно сьогодні вночі, щоб погасити свій сонний борг.

- Включайте години сьогоднішньої ночі у загальний час, необхідний для відновлення.
- Якщо ви вже виспалися достатньо, щоб покрити ціль на сьогоднішню ніч або більше, поверніть `0`.

# --hints--

`sleepDebt([6, 6, 6, 6, 6, 6], 8)` має повертати `20`.

```js
assert.equal(sleepDebt([6, 6, 6, 6, 6, 6], 8), 20);
```

`sleepDebt([6, 7, 8, 4, 8, 6], 7)` має повертати `10`.

```js
assert.equal(sleepDebt([6, 7, 8, 4, 8, 6], 7), 10);
```

`sleepDebt([10, 10, 9, 10, 9, 11], 9)` має повертати `4`.

```js
assert.equal(sleepDebt([10, 10, 9, 10, 9, 11], 9), 4);
```

`sleepDebt([8, 7, 6, 7, 6, 8], 6)` має повертати `0`.

```js
assert.equal(sleepDebt([8, 7, 6, 7, 6, 8], 6), 0);
```

`sleepDebt([8, 9, 10, 9, 10, 7], 7)` має повертати `0`.

```js
assert.equal(sleepDebt([8, 9, 10, 9, 10, 7], 7), 0);
```

# --seed--

## --seed-contents--

```js
function sleepDebt(hoursSlept, targetHours) {

  return hoursSlept;
}
```

# --solutions--

```js
function sleepDebt(hoursSlept, targetHours) {
  const debt = (hoursSlept.length + 1) * targetHours - hoursSlept.reduce((sum, h) => sum + h, 0);
  return Math.max(0, debt);
}
```
