---
id: 69f35a5bb823ed620fcb7cb8
title: "Завдання 279: найдовший ланцюжок доміно"
challengeType: 28
dashedName: challenge-279
---

# --description--

Дано двовимірний масив, що представляє множину доміно, поверніть найдовший дійсний ланцюжок.

- Кожне доміно — це пара чисел від 0 до 6, наприклад `[3, 2]`.
- Ланцюжок є дійсним, коли друге число кожного доміно збігається з першим числом наступного.
- Перше число першого доміно та друге число останнього не повинні співпадати з чимось.
- Будь-яке доміно можна перевернути, тому `[3, 2]` можна зіграти як `[2, 3]`.
- Завжди існує рівно один найдовший дійсний ланцюжок.

Наприклад, для `[[1, 2], [4, 5], [2, 3]]` поверніть `[[1, 2], [2, 3]]`.

# --hints--

`getLongestChain([[1, 2], [4, 5], [2, 3]])` має повернути `[[1, 2], [2, 3]]`.

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [4, 5], [2, 3]]));
const result1 = JSON.stringify([[1, 2], [2, 3]]);
const result2 = JSON.stringify([[3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[2, 1], [4, 3], [5, 3]])` має повернути `[[4, 3], [3, 5]]`.

```js
const chain = JSON.stringify(getLongestChain([[2, 1], [4, 3], [5, 3]]));
const result1 = JSON.stringify([[4, 3], [3, 5]]);
const result2 = JSON.stringify([[5, 3], [3, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]])` має повернути `[[1, 2], [2, 3], [3, 4], [4, 0]]`.

```js
const chain = JSON.stringify(getLongestChain([[1, 2], [3, 4], [2, 3], [4, 0]]));
const result1 = JSON.stringify([[1, 2], [2, 3], [3, 4], [4, 0]]);
const result2 = JSON.stringify([[0, 4], [4, 3], [3, 2], [2, 1]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]])` має повернути `[[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]`.

```js
const chain = JSON.stringify(getLongestChain([[6, 6], [6, 1], [1, 1], [0, 3], [2, 3], [4, 1], [5, 6]]));
const result1 = JSON.stringify([[4, 1], [1, 1], [1, 6], [6, 6], [6, 5]]);
const result2 = JSON.stringify([[5, 6], [6, 6], [6, 1], [1, 1], [1, 4]]);
assert.isTrue(chain === result1 || chain === result2);
```

`getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]])` має повернути `[[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]`.

```js
const chain = JSON.stringify(getLongestChain([[0, 4], [3, 3], [0, 3], [5, 6], [4, 5], [4, 2], [5, 5], [1, 2], [4, 4]]));
const result1 = JSON.stringify([[3, 3], [3, 0], [0, 4], [4, 4], [4, 5], [5, 5], [5, 6]]);
const result2 = JSON.stringify([[6, 5], [5, 5], [5, 4], [4, 4], [4, 0], [0, 3], [3, 3]]);
assert.isTrue(chain === result1 || chain === result2);
```

# --seed--

## --seed-contents--

```js
function getLongestChain(dominoes) {

  return dominoes;
}
```

# --solutions--

```js
function getLongestChain(dominoes) {
  function search(chain, remaining) {
    let best = chain;
    const last = chain[chain.length - 1][1];

    for (let i = 0; i < remaining.length; i++) {
      const [a, b] = remaining[i];
      const rest = remaining.filter((_, j) => j !== i);

      if (a === last) {
        const result = search([...chain, [a, b]], rest);
        if (result.length > best.length) best = result;
      }
      if (b === last && a !== b) {
        const result = search([...chain, [b, a]], rest);
        if (result.length > best.length) best = result;
      }
    }
    return best;
  }

  let best = [];
  for (let i = 0; i < dominoes.length; i++) {
    const [a, b] = dominoes[i];
    const rest = dominoes.filter((_, j) => j !== i);

    const r1 = search([[a, b]], rest);
    if (r1.length > best.length) best = r1;

    if (a !== b) {
      const r2 = search([[b, a]], rest);
      if (r2.length > best.length) best = r2;
    }
  }
  return best;
}
```
