---
id: 69f35a5bb823ed620fcb7cbd
title: "Завдання 284: I перед E"
challengeType: 28
dashedName: challenge-284
---

# --description--

Дано слово або речення, поверніть виправлену версію, де кожне слово відповідає правилу «I перед E, крім випадків після C».

- Якщо слово містить `"ei"`, яке не передує `"c"`, замініть його на `"ie"`.
- Якщо слово містить `"ie"`, перед яким стоїть `"c"`, замініть його на `"ei"`.
- Усі інші слова залишаються без змін.

# --hints--

`iBeforeE("beleive")` має повернути `"believe"`.

```js
assert.equal(iBeforeE("beleive"), "believe");
```

`iBeforeE("recieve")` має повернути `"receive"`.

```js
assert.equal(iBeforeE("recieve"), "receive");
```

`iBeforeE("we recieved a breif")` має повернути `"we received a brief"`.

```js
assert.equal(iBeforeE("we recieved a breif"), "we received a brief");
```

`iBeforeE("she beleived the friendly niece could percieve the greif")` має повернути `"she believed the friendly niece could perceive the grief"`.

```js
assert.equal(iBeforeE("she beleived the friendly niece could percieve the greif"), "she believed the friendly niece could perceive the grief");
```

`iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit")` має повернути `"we received relief after the thief gave us a brief piece of fierce deceit"`.

```js
assert.equal(iBeforeE("we recieved relief after the theif gave us a breif piece of feirce deceit"), "we received relief after the thief gave us a brief piece of fierce deceit");
```

# --seed--

## --seed-contents--

```js
function iBeforeE(sentence) {

  return sentence;
}
```

# --solutions--

```js
function iBeforeE(sentence) {
  return sentence.split(' ').map(word => {
    word = word.replace(/([^c])ei/g, '$1ie');
    word = word.replace(/cie/g, 'cei');
    return word;
  }).join(' ');
}
```
