---
id: 69f35a5bb823ed620fcb7cbc
title: "Завдання 283: рядковий зіпер"
challengeType: 28
dashedName: challenge-283
---

# --description--

Дано два рядки, поверніть новий рядок, який по черзі чергує символи з обох рядків по одному. Якщо один рядок довший, додайте залишкові символи в кінці.

Починайте з першого символу першого рядка.

# --hints--

`zipStrings("abc", "123")` має повертати `"a1b2c3"`.

```js
assert.equal(zipStrings("abc", "123"), "a1b2c3");
```

`zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz")` має повертати `"abcdefghijklmnopqrstuvwxyz"`.

```js
assert.equal(zipStrings("acegikmoqsuwy", "bdfhjlnprtvxz"), "abcdefghijklmnopqrstuvwxyz");
```

`zipStrings("day", "night")` має повертати `"dnaiyght"`.

```js
assert.equal(zipStrings("day", "night"), "dnaiyght");
```

`zipStrings("python", "javascript")` має повертати `"pjyatvhaosncript"`.

```js
assert.equal(zipStrings("python", "javascript"), "pjyatvhaosncript");
```

`zipStrings("feCdCm", "reoeap")` має повертати `"freeCodeCamp"`.

```js
assert.equal(zipStrings("feCdCm", "reoeap"), "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
function zipStrings(a, b) {

  return a;
}
```

# --solutions--

```js
function zipStrings(a, b) {
  let result = '';
  const len = Math.max(a.length, b.length);
  for (let i = 0; i < len; i++) {
    if (i < a.length) result += a[i];
    if (i < b.length) result += b[i];
  }
  return result;
}
```
