---
id: 69f8c998d78ad3171a0713bb
title: "Завдання 288: секретне число"
challengeType: 28
dashedName: challenge-288
---

# --description--

За заданим секретним числом і припущенням визначте, чи є припущення правильним.

Поверніть:

- `"higher"`, якщо секретне число більше за припущення.
- `"lower"`, якщо секретне число менше за припущення.
- `"you got it!"`, якщо припущення правильне.

# --hints--

`guessNumber(50, 30)` має повертати `"higher"`.

```js
assert.equal(guessNumber(50, 30), "higher");
```

`guessNumber(85, 99)` має повертати `"lower"`.

```js
assert.equal(guessNumber(85, 99), "lower");
```

`guessNumber(2026, 2026)` має повертати `"you got it!"`.

```js
assert.equal(guessNumber(2026, 2026), "you got it!");
```

`guessNumber(92904, 11283)` має повертати `"higher"`.

```js
assert.equal(guessNumber(92904, 11283), "higher");
```

`guessNumber(230495, 423920)` має повертати `"lower"`.

```js
assert.equal(guessNumber(230495, 423920), "lower");
```

`guessNumber(120349, 120349)` має повертати `"you got it!"`.

```js
assert.equal(guessNumber(120349, 120349), "you got it!");
```

# --seed--

## --seed-contents--

```js
function guessNumber(secret, guess) {

  return secret;
}
```

# --solutions--

```js
function guessNumber(secret, guess) {
  if (guess === secret) return "you got it!";
  return guess < secret ? "higher" : "lower";
}
```
