---
id: 69f8c998d78ad3171a0713be
title: "Завдання 291: підрахунок fizzbuzz"
challengeType: 28
dashedName: challenge-291
---

# --description--

За заданими початковим і кінцевим числом порахуйте кількість появ fizz і buzz у діапазоні (включно).

- Числа, кратні 3, рахуйте як fizz.
- Числа, кратні 5, рахуйте як buzz.
- Числа, кратні і 3, і 5, рахуйте як одночасно fizz і buzz.

Поверніть об’єкт або словник із підрахунками у форматі: `{ fizz, buzz }`.

# --hints--

`fizzBuzzCount(1, 11)` має повертати `{fizz: 3, buzz: 2}`.

```js
assert.deepEqual(fizzBuzzCount(1, 11), {fizz: 3, buzz: 2});
```

`fizzBuzzCount(14, 41)` має повертати `{fizz: 9, buzz: 6}`.

```js
assert.deepEqual(fizzBuzzCount(14, 41), {fizz: 9, buzz: 6});
```

`fizzBuzzCount(24, 100)` має повертати `{fizz: 26, buzz: 16}`.

```js
assert.deepEqual(fizzBuzzCount(24, 100), {fizz: 26, buzz: 16});
```

`fizzBuzzCount(-635, -14)` має повертати `{fizz: 207, buzz: 125}`.

```js
assert.deepEqual(fizzBuzzCount(-635, -14), {fizz: 207, buzz: 125});
```

`fizzBuzzCount(-5432, 6789)` має повертати `{fizz: 4074, buzz: 2444}`.

```js
assert.deepEqual(fizzBuzzCount(-5432, 6789), {fizz: 4074, buzz: 2444});
```

# --seed--

## --seed-contents--

```js
function fizzBuzzCount(start, end) {

  return start;
}
```

# --solutions--

```js
function fizzBuzzCount(start, end) {
  let fizz = 0, buzz = 0;

  for (let i = start; i <= end; i++) {
    if (i % 3 === 0) fizz++;
    if (i % 5 === 0) buzz++;
  }

  return { fizz, buzz };
}
```
