---
id: 69f8c998d78ad3171a0713b9
title: "Завдання 286: відкриті проблеми"
challengeType: 28
dashedName: challenge-286
---

# --description--

Дано масив номерів проблем і ще один масив номерів запитів на злиття (PR). Поверніть масив проблем, які залишаються відкритими після того, як усі PR були об’єднані.

- PR закриває проблему, якщо їхні цифри є обертанням одна одної. Наприклад, проблему 123 закриє PR 231 або 312.
- PR не закриває проблему з точно таким самим числом. Наприклад, PR 123 не закриває проблему 123. Отже, проблему з усіма однаковими цифрами не можна закрити.
- Будь-яке число може мати видалені провідні нулі. Наприклад, PR 201 закриє проблему 12 (012, обертання 201). Аналогічно, проблему 201 закриє PR 12.

Поверніть залишкові відкриті проблеми у тому порядку, в якому їх було задано.

# --hints--

`getOpenIssues([123, 234], [231])` має повертати `[234]`.

```js
assert.deepEqual(getOpenIssues([123, 234], [231]), [234]);
```

`getOpenIssues([123, 345, 16], [345, 231])` має повертати `[345, 16]`.

```js
assert.deepEqual(getOpenIssues([123, 345, 16], [345, 231]), [345, 16]);
```

`getOpenIssues([456, 332, 12, 15], [201, 945, 180])` має повертати `[456, 332, 15]`.

```js
assert.deepEqual(getOpenIssues([456, 332, 12, 15], [201, 945, 180]), [456, 332, 15]);
```

`getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21])` має повертати `[115, 296, 24]`.

```js
assert.deepEqual(getOpenIssues([12, 115, 296, 170, 24], [17, 18, 19, 20, 21]), [115, 296, 24]);
```

`getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14])` має повертати `[95, 395, 754, 296, 709, 237, 1802]`.

```js
assert.deepEqual(getOpenIssues([19, 95, 422, 395, 754, 102, 296, 709, 237, 4400, 1802], [395, 440, 9001, 95, 242, 21, 287, 169, 14]), [95, 395, 754, 296, 709, 237, 1802]);
```

# --seed--

## --seed-contents--

```js
function getOpenIssues(issues, prs) {

  return issues;
}
```

# --solutions--

```js
function getOpenIssues(issues, prs) {
  function getRotations(n) {
    const s = String(n);
    const rotations = new Set();
    for (let i = 0; i < s.length; i++) {
      const rotated = s.slice(i) + s.slice(0, i);
      rotations.add(Number(rotated));
    }
    return rotations;
  }

  function areRotations(a, b) {
    if (a === b) return false;
    const rotations = getRotations(a);
    return rotations.has(b);
  }

  return issues.filter(issue =>
    !prs.some(pr => areRotations(pr, issue) || areRotations(issue, pr))
  );
}
```
