---
id: 69f8c998d78ad3171a0713ba
title: "Завдання 287: виправлення римських чисел"
challengeType: 28
dashedName: challenge-287
---

# --description--

Дано рядок з неправильно записаними римськими числами, поверніть значення у стандартній римській нотації.

Вхідні дані використовуватимуть лише адитивну нотацію, тому кожен символ додає своє значення до загальної суми. Нагадаємо символи та їхні значення:

| Символ | Значення |
| ------ | -------- |
| `"I"`  | 1        |
| `"V"`  | 5        |
| `"X"`  | 10       |
| `"L"`  | 50       |
| `"C"`  | 100      |
| `"D"`  | 500      |
| `"M"`  | 1000     |

При повторному кодуванні використовуйте найбільший можливий символ на кожному кроці, застосовуючи віднімальні пари (`"IV"`, `"IX"`, `"XL"`, `"XC"`, `"CD"`, `"CM"`) там, де це потрібно.

# --hints--

`fixNumerals("XIIIII")` має повернути `"XV"`.

```js
assert.equal(fixNumerals("XIIIII"), "XV");
```

`fixNumerals("IIIILX")` має повернути `"LXIV"`.

```js
assert.equal(fixNumerals("IIIILX"), "LXIV");
```

`fixNumerals("XXVVVIIIII")` має повернути `"XL"`.

```js
assert.equal(fixNumerals("XXVVVIIIII"), "XL");
```

`fixNumerals("MDCCLXXXXVIIII")` має повернути `"MDCCXCIX"`.

```js
assert.equal(fixNumerals("MDCCLXXXXVIIII"), "MDCCXCIX");
```

`fixNumerals("IIIIVVVVXXXXLLLLCCDD")` має повернути `"MCDLXIV"`.

```js
assert.equal(fixNumerals("IIIIVVVVXXXXLLLLCCDD"), "MCDLXIV");
```

`fixNumerals("ILCDMIVDIIXLCVCXDL")` має повернути `"MMCMLXXXIV"`.

```js
assert.equal(fixNumerals("ILCDMIVDIIXLCVCXDL"), "MMCMLXXXIV");
```

# --seed--

## --seed-contents--

```js
function fixNumerals(str) {

  return str;
}
```

# --solutions--

```js
function fixNumerals(str) {
  const values = { I: 1, V: 5, X: 10, L: 50, C: 100, D: 500, M: 1000 };
  const encoding = [
    [1000, "M"], [900, "CM"], [500, "D"], [400, "CD"],
    [100, "C"], [90, "XC"], [50, "L"], [40, "XL"],
    [10, "X"], [9, "IX"], [5, "V"], [4, "IV"], [1, "I"]
  ];

  let total = [...str].reduce((sum, c) => sum + values[c], 0);

  let result = "";
  for (const [value, symbol] of encoding) {
    while (total >= value) {
      result += symbol;
      total -= value;
    }
  }
  return result;
}
```
