---
id: 69f8c998d78ad3171a0713bf
title: "Завдання 292: ширше співвідношення сторін"
challengeType: 28
dashedName: challenge-292
---

# --description--

Дано два рядки з різними розмірами зображень, поверніть співвідношення сторін зображення з більшим співвідношенням ширини до висоти.

- Дано рядки у форматі `"WxH"`, наприклад, `"1920x1080"`.
- Співвідношення сторін — це відношення ширини до висоти, скорочене до найменших цілих чисел. Наприклад, `"1920x1080"` скорочується до `"16:9"`.
- Поверніть рядок у форматі `"W:H"`, наприклад, `"16:9"`.

# --hints--

`getWiderAspectRatio("1920x1080", "800x600")` має повернути `"16:9"`.

```js
assert.equal(getWiderAspectRatio("1920x1080", "800x600"), "16:9");
```

`getWiderAspectRatio("1080x1350", "2048x1536")` має повернути `"4:3"`.

```js
assert.equal(getWiderAspectRatio("1080x1350", "2048x1536"), "4:3");
```

`getWiderAspectRatio("640x480", "2440x1220")` має повернути `"2:1"`.

```js
assert.equal(getWiderAspectRatio("640x480", "2440x1220"), "2:1");
```

`getWiderAspectRatio("360x640", "1080x1920")` має повернути `"9:16"`.

```js
assert.equal(getWiderAspectRatio("360x640", "1080x1920"), "9:16");
```

`getWiderAspectRatio("3440x1440", "2048x858")` має повернути `"43:18"`.

```js
assert.equal(getWiderAspectRatio("3440x1440", "2048x858"), "43:18");
```

`getWiderAspectRatio("12345x61234", "12534x51234")` має повернути `"2089:8539"`.

```js
assert.equal(getWiderAspectRatio("12345x61234", "12534x51234"), "2089:8539");
```

# --seed--

## --seed-contents--

```js
function getWiderAspectRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function getWiderAspectRatio(a, b) {
  function gcd(x, y) {
    return y === 0 ? x : gcd(y, x % y);
  }

  function parse(str) {
    const [w, h] = str.split("x").map(Number);
    const g = gcd(w, h);
    return { ratio: w / h, simplified: `${w / g}:${h / g}` };
  }

  const imageA = parse(a);
  const imageB = parse(b);

  return imageA.ratio >= imageB.ratio ? imageA.simplified : imageB.simplified;
}
```
