---
id: 69f8c998d78ad3171a0713bd
title: "Завдання 290: піца на вечірку"
challengeType: 28
dashedName: challenge-290
---

# --description--

Дано масив годин, відпрацьованих сьогодні кожною людиною. Потрібно повернути кількість піц, які слід замовити для піцерійної вечірки.

- Розділіть кількість годин кожної людини на 3, щоб отримати кількість шматочків піци.
- Ви не можете з’їсти частковий шматок, тому округліть кількість шматочків кожної людини вгору до найближчого цілого числа.
- Кожна людина отримує мінімум два шматочки.
- Кожна піца має 8 шматочків. Округліть загальну кількість піц вгору до найближчої цілої піци.

# --hints--

`getPizzasToOrder([8, 8, 8])` має повернути `2`.

```js
assert.equal(getPizzasToOrder([8, 8, 8]), 2);
```

`getPizzasToOrder([10, 9, 8, 2, 2, 6, 10])` має повернути `3`.

```js
assert.equal(getPizzasToOrder([10, 9, 8, 2, 2, 6, 10]), 3);
```

`getPizzasToOrder([1, 2, 3, 4, 5])` має повернути `2`.

```js
assert.equal(getPizzasToOrder([1, 2, 3, 4, 5]), 2);
```

`getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8])` має повернути `3`.

```js
assert.equal(getPizzasToOrder([8, 8, 8, 8, 8, 8, 8, 8]), 3);
```

`getPizzasToOrder([9, 9, 6])` має повернути `1`.

```js
assert.equal(getPizzasToOrder([9, 9, 6]), 1);
```

`getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0])` має повернути `5`.

```js
assert.equal(getPizzasToOrder([10, 12, 16, 9, 8, 11, 15, 8, 0]), 5);
```

# --seed--

## --seed-contents--

```js
function getPizzasToOrder(hoursWorked) {

  return hoursWorked;
}
```

# --solutions--

```js
function getPizzasToOrder(hoursWorked) {
  const totalSlices = hoursWorked.reduce((sum, hours) => {
    return sum + Math.max(Math.ceil(hours / 3), 2);
  }, 0);

  return Math.ceil(totalSlices / 8);
}
```
