---
id: 69f90c1329a94b37e2a2086d
title: "Завдання 294: комбінації круглих дужок"
challengeType: 28
dashedName: challenge-294
---

# --description--

Дано ціле число `n`, поверніть кількість дійсних комбінацій `n` пар круглих дужок.

- Дійсна комбінація — це рядок, у якому кожна відкрита кругла дужка має відповідну закриту круглу дужку, і жодна закрита кругла дужка не з’являється раніше за свою відповідну відкриту.

Наприклад, для `2` існує 2 дійсні комбінації:

```json
(())
()()
```

# --hints--

`getCombinations(2)` має повернути `2`.

```js
assert.equal(getCombinations(2), 2);
```

`getCombinations(3)` має повернути `5`.

```js
assert.equal(getCombinations(3), 5);
```

`getCombinations(5)` має повернути `42`.

```js
assert.equal(getCombinations(5), 42);
```

`getCombinations(8)` має повернути `1430`.

```js
assert.equal(getCombinations(8), 1430);
```

`getCombinations(13)` має повернути `742900`.

```js
assert.equal(getCombinations(13), 742900);
```

# --seed--

## --seed-contents--

```js
function getCombinations(n) {

  return n;
}
```

# --solutions--

```js
function getCombinations(n) {
  let count = 0;

  function generate(open, close) {
    if (open === 0 && close === 0) {
      count++;
      return;
    }
    if (open > 0) generate(open - 1, close);
    if (close > open) generate(open, close - 1);
  }

  generate(n, n);
  return count;
}
```
