---
id: 6a0dcc730cb92a616f86f0c1
title: "Завдання 297: Валідатор схеми. Частина 3"
challengeType: 28
dashedName: challenge-297
---

# --description--

Дано об’єкт (JavaScript) або словник (Python). Визначте, чи відповідає він наступній схемі:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

{
  username: string,
  posts: number,
  verified: boolean,
  role: Roles
}
```

- Символ вертикальної риски (`|`) означає "або". `role` має бути одним із перелічених `Roles` значень.
- Додаткові ключі дозволені

# --hints--

`isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ username: "henry", posts: 0, verified: true, role: "staff" }));
```

`isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ username: "sara", posts: 45, verified: false, role: "creator", followers: 70 }));
```

`isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ username: "penelope", posts: 20, verified: true, role: "admin" }));
```

`isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ username: "kevin", posts: 0, verified: false, role: "user" }));
```

`isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ username: "george", posts: 15, verified: true, role: "moderator" }));
```

`isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: "david", posts: 0, verified: false, role: "guest" }));
```

`isValidSchema({ username: "wendy", posts: 10, verified: true })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: "wendy", posts: 10, verified: true }));
```

`isValidSchema({ username: "fabian", posts: 1, verified: true, role: true })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: "fabian", posts: 1, verified: true, role: true }));
```

`isValidSchema({ username: 8, posts: 1, verified: true, role: "user" })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: 8, posts: 1, verified: true, role: "user" }));
```

`isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: "penny", posts: "10", verified: true, role: "staff" }));
```

`isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" })` має повертати `false`.

```js
assert.isFalse(isValidSchema({ username: "john", posts: "1", verified: "true", role: "admin" }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];
  return (
    typeof obj.username === 'string' &&
    typeof obj.posts === 'number' &&
    typeof obj.verified === 'boolean' &&
    roles.includes(obj.role)
  );
}
```
