---
id: 6a0dcd03ee4e68698080ef68
title: "Завдання 304: організація маршруту"
challengeType: 28
dashedName: challenge-304
---

# --description--

Дано масив щонайменше з двома необов’язковими зупинками для одноденної поїздки. Потрібно повернути кількість дійсних варіантів організації маршруту.

Маршрут завжди включає `"breakfast"`, `"lunch"` і `"dinner"`, їх не передають як аргументи. Необов’язкові зупинки можна розміщувати будь-де в маршруті за такими правилами:

- `"breakfast"` завжди перший, перед `"lunch"` має бути щонайменше одна зупинка.
- `"lunch"` має з’являтися перед `"dinner"`, між ними має бути щонайменше одна зупинка.
- Після `"dinner"` може бути не більше однієї необов’язкової зупинки.

Поверніть кількість дійсних варіантів організації.

# --hints--

`getItineraryCount(["library", "park"])` має повертати `2`.

```js
assert.equal(getItineraryCount(["library", "park"]), 2);
```

`getItineraryCount(["library", "park", "arcade"])` має повертати `18`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade"]), 18);
```

`getItineraryCount(["library", "park", "arcade", "store"])` має повертати `120`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store"]), 120);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe"])` має повертати `840`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe"]), 840);
```

`getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"])` має повертати `55440`.

```js
assert.equal(getItineraryCount(["library", "park", "arcade", "store", "cafe", "market", "museum"]), 55440);
```

# --seed--

## --seed-contents--

```js
function getItineraryCount(stops) {

  return stops;
}
```

# --solutions--

```js
function getItineraryCount(stops) {
  const items = ["lunch", "dinner", ...stops];
  let count = 0;

  function permute(arr) {
    if (arr.length === 0) return [[]];
    return arr.flatMap((item, i) =>
      permute([...arr.slice(0, i), ...arr.slice(i + 1)]).map(p => [item, ...p])
    );
  }

  for (const perm of permute(items)) {
    const lunchIdx = perm.indexOf("lunch");
    const dinnerIdx = perm.indexOf("dinner");

    if (lunchIdx >= dinnerIdx) continue;
    if (lunchIdx < 1) continue;
    if (dinnerIdx - lunchIdx < 2) continue;
    if (perm.length - dinnerIdx - 1 > 1) continue;

    count++;
  }

  return count;
}
```
