---
id: 6a0dcd03ee4e68698080ef66
title: "Завдання 302: Джетлаг"
challengeType: 28
dashedName: challenge-302
---

# --description--

За заданим містом відправлення, містом прибуття, тривалістю польоту в годинах і напрямком подорожі поверніть кількість годин джетлагу, які відчуває мандрівник.

Задані міста будуть із наступного списку з їхнім UTC-компенсуванням:

| Місто | Компенсування |
| - | - |
| `"Los Angeles"` | -8 |
| `"New York"` | -5 |
| `"London"` | 0 |
| `"Istanbul"` | +3 |
| `"Dubai"` | +4 |
| `"Hong Kong"` | +8 |
| `"Tokyo"` | +9 |

Щоб обчислити години джетлагу:

1. Знайдіть різницю в часових поясах у годинах між двома містами.
2. Визначте множник напрямку. Якщо подорож відбувається `"east"`, він дорівнює 1.5, інакше — 1.0.
3. Отримайте години джетлагу за формулою: різниця в часових поясах + (тривалість польоту * 0.1) * множник напрямку

Поверніть години джетлагу, округлені до одного десяткового знака.

# --hints--

`getJetLagHours("Istanbul", "Hong Kong", 10, "east")` має повернути `6.5`.

```js
assert.equal(getJetLagHours("Istanbul", "Hong Kong", 10, "east"), 6.5);
```

`getJetLagHours("London", "New York", 8, "west")` має повернути `5.8`.

```js
assert.equal(getJetLagHours("London", "New York", 8, "west"), 5.8);
```

`getJetLagHours("Hong Kong", "Tokyo", 4, "east")` має повернути `1.6`.

```js
assert.equal(getJetLagHours("Hong Kong", "Tokyo", 4, "east"), 1.6);
```

`getJetLagHours("Dubai", "London", 7, "west")` має повернути `4.7`.

```js
assert.equal(getJetLagHours("Dubai", "London", 7, "west"), 4.7);
```

`getJetLagHours("Los Angeles", "Hong Kong", 15, "west")` має повернути `17.5`.

```js
assert.equal(getJetLagHours("Los Angeles", "Hong Kong", 15, "west"), 17.5);
```

`getJetLagHours("Tokyo", "Dubai", 9, "west")` має повернути `5.9`.

```js
assert.equal(getJetLagHours("Tokyo", "Dubai", 9, "west"), 5.9);
```

`getJetLagHours("New York", "Istanbul", 10, "east")` має повернути `9.5`.

```js
assert.equal(getJetLagHours("New York", "Istanbul", 10, "east"), 9.5);
```

# --seed--

## --seed-contents--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {

  return departureCity;
}
```

# --solutions--

```js
function getJetLagHours(departureCity, arrivalCity, flightDuration, direction) {
  const offsets = {
    "Los Angeles": -8,
    "New York": -5,
    "London": 0,
    "Istanbul": +3,
    "Dubai": +4,
    "Hong Kong": +8,
    "Tokyo": +9
  };

  const depOffset = offsets[departureCity];
  const arrOffset = offsets[arrivalCity];

  const timezoneDiff = Math.abs(arrOffset - depOffset);
  const multiplier = direction === "east" ? 1.5 : 1.0;

  const score = timezoneDiff + (flightDuration * 0.1) * multiplier;
  return Math.round(score * 10) / 10;
}
```
