---
id: 6a0dcc730cb92a616f86f0bf
title: "Завдання 295: Валідатор схеми. Частина 1"
challengeType: 28
dashedName: challenge-295
---

# --description--

Дано об’єкт (JavaScript) або словник (Python). Визначте, чи відповідає він наступній схемі:

```json
{
  username: string
}
```

- Додаткові ключі дозволені

# --hints--

`isValidSchema({ username: "bob" })` має повернути `true`.

```js
assert.isTrue(isValidSchema({ username: "bob" }));
```

`isValidSchema({ username: "jen", posts: 30 })` має повернути `true`.

```js
assert.isTrue(isValidSchema({ username: "jen", posts: 30 }));
```

`isValidSchema({ username: "" })` має повернути `true`.

```js
assert.isTrue(isValidSchema({ username: "" }));
```

`isValidSchema({ username: 7 })` має повернути `false`.

```js
assert.isFalse(isValidSchema({ username: 7 }));
```

`isValidSchema({ posts: 25 })` має повернути `false`.

```js
assert.isFalse(isValidSchema({ posts: 25 }));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  return typeof obj.username === 'string';
}
```
