---
id: 6a0dcd03ee4e68698080ef6c
title: "Завдання 308: перевірка кредитної картки"
challengeType: 28
dashedName: challenge-308
---

# --description--

Дано рядок цифр, що представляє номер кредитної картки. Визначте, чи є цей номер дійсним, використовуючи такий метод:

- Починаючи з передостанньої цифри, подвоюйте кожну другу цифру, рухаючись вліво.
- Якщо подвоєння цифри дає число більше за 9, відніміть 9.
- Підсумуйте всі цифри (подвоєні та неподвоєні).
- Якщо сума ділиться на 10 без остачі, номер є дійсним.

# --hints--

`isValidCard("4532015112830366")` має повертати `true`.

```js
assert.isTrue(isValidCard("4532015112830366"));
```

`isValidCard("5425233430109903")` має повертати `true`.

```js
assert.isTrue(isValidCard("5425233430109903"));
```

`isValidCard("371449635398431")` має повертати `true`.

```js
assert.isTrue(isValidCard("371449635398431"));
```

`isValidCard("6011111111111117")` має повертати `true`.

```js
assert.isTrue(isValidCard("6011111111111117"));
```

`isValidCard("4532015112830367")` має повертати `false`.

```js
assert.isFalse(isValidCard("4532015112830367"));
```

`isValidCard("1234567890123456")` має повертати `false`.

```js
assert.isFalse(isValidCard("1234567890123456"));
```

`isValidCard("4532015112830368")` має повертати `false`.

```js
assert.isFalse(isValidCard("4532015112830368"));
```

# --seed--

## --seed-contents--

```js
function isValidCard(number) {

  return number;
}
```

# --solutions--

```js
function isValidCard(number) {
  let sum = 0;

  for (let i = number.length - 1; i >= 0; i--) {
    let digit = parseInt(number[i]);

    if ((number.length - 1 - i) % 2 === 1) {
      digit *= 2;
      if (digit > 9) digit -= 9;
    }

    sum += digit;
  }

  return sum % 10 === 0;
}
```
