---
id: 6a0dcd03ee4e68698080ef6b
title: "Завдання 307: правила зонування"
challengeType: 28
dashedName: challenge-307
---

# --description--

Дано двовимірну сітку (масив масивів), що представляє макет будівель міста. Поверніть координати всіх будівель, які порушують правила зонування.

Кожна клітинка сітки містить одну з міток із таблиці нижче. Будівля порушує правила, якщо будь-який із її (до) 4 сусідів по горизонталі або вертикалі є типом, з яким вона не може бути поруч.

| Мітка | Тип | Не може бути поруч із |
| - | - | - |
| `"i"` | промислова | `"R"`, `"I"` |
| `"A"` | сільськогосподарська | `"C"` |
| `"R"` | житлова | `"i"`, `"C"` |
| `"I"` | інституційна | `"i"` |
| `"C"` | комерційна | `"R"`, `"A"` |
| `""` (порожній рядок) | незабудована | без обмежень |

Поверніть координати всіх клітинок, що порушують правила, як масив пар `[row, col]`, у будь-якому порядку. Якщо порушень немає, поверніть порожній масив.

# --hints--

`getZoneViolations([["R", "C"], ["", "C"]])` має повертати `[[0, 0], [0, 1]]`.

```js
assert.deepEqual(getZoneViolations([["R", "C"], ["", "C"]]).sort(), [[0, 0], [0, 1]]);
```

`getZoneViolations([["", "i"], ["", "R"], ["R", "I"]])` має повертати `[[0, 1], [1, 1]]`.

```js
assert.deepEqual(getZoneViolations([["", "i"], ["", "R"], ["R", "I"]]).sort(), [[0, 1], [1, 1]]);
```

`getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]])` має повертати `[]`.

```js
assert.deepEqual(getZoneViolations([["A", "i", "C"], ["A", "", "C"], ["R", "R", "I"]]).sort(), []);
```

`getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]])` має повертати `[[0, 1], [0, 2], [0, 3]]`.

```js
assert.deepEqual(getZoneViolations([["R", "R", "C", "R", "R"], ["R", "I", "C", "", "A"], ["R", "R", "", "i", "A"]]).sort(), [[0, 1], [0, 2], [0, 3]]);
```


`getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]])` має повертати `[[2, 3], [2, 4], [3, 1], [3, 2]]`.

```js
assert.deepEqual(getZoneViolations([["R", "A", "A", "", "i", "i"], ["R", "I", "", "C", "i", "i"], ["R", "", "C", "C", "A", "A"], ["R", "R", "C", "I", "R", "R"]]).sort(), [[2, 3], [2, 4], [3, 1], [3, 2]]);
```

# --seed--

## --seed-contents--

```js
function getZoneViolations(grid) {

  return grid;
}
```

# --solutions--

```js
function getZoneViolations(grid) {
  const rules = {
    "i": ["R", "I"],
    "A": ["C"],
    "R": ["i", "C"],
    "I": ["i"],
    "C": ["R", "A"],
    "": []
  };

  const directions = [[-1, 0], [0, -1], [0, 1], [1, 0]];
  const violations = [];

  for (let row = 0; row < grid.length; row++) {
    for (let col = 0; col < grid[row].length; col++) {
      const cell = grid[row][col];
      const forbidden = rules[cell];

      for (const [dr, dc] of directions) {
        const nr = row + dr;
        const nc = col + dc;
        if (nr >= 0 && nr < grid.length && nc >= 0 && nc < grid[row].length) {
          if (forbidden.includes(grid[nr][nc])) {
            violations.push([row, col]);
            break;
          }
        }
      }
    }
  }

  return violations;
}
```
