---
id: 6a0dcc730cb92a616f86f0c5
title: "Завдання 301: останнє завантаження"
challengeType: 28
dashedName: challenge-301
---

# --description--

Враховуючи кількість мірок прального порошку, що залишилися, та масив із кількістю мірок, які ви використовували в кожен із попередніх днів, поверніть кількість повних днів, на які вистачить порошку.

Обчисліть середнє щоденне використання за історією використання і припустіть, що саме таку кількість ви використовуватимете щодня надалі.

# --hints--

`lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2])` має повертати `5`.

```js
assert.equal(lastLoadDate(10, [2, 2, 2, 2, 2, 2, 2]), 5);
```

`lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1])` має повертати `7`.

```js
assert.equal(lastLoadDate(16, [2, 3, 0, 3, 4, 2, 1]), 7);
```

`lastLoadDate(33, [5, 0, 4, 3, 3, 2])` має повертати `11`.

```js
assert.equal(lastLoadDate(33, [5, 0, 4, 3, 3, 2]), 11);
```

`lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2])` має повертати `12`.

```js
assert.equal(lastLoadDate(50, [2, 0, 2, 9, 12, 0, 2]), 12);
```

`lastLoadDate(20, [13, 9, 12, 10, 8])` має повертати `1`.

```js
assert.equal(lastLoadDate(20, [13, 9, 12, 10, 8]), 1);
```

# --seed--

## --seed-contents--

```js
function lastLoadDate(scoops, usage) {

  return scoops;
}
```

# --solutions--

```js
function lastLoadDate(scoops, usage) {
  const avg = usage.reduce((sum, n) => sum + n, 0) / usage.length;
  return Math.floor(scoops / avg);
}
```
