---
id: 6a0dcc730cb92a616f86f0c4
title: "Завдання 300: Валідатор схеми Частина 6"
challengeType: 28
dashedName: challenge-300
---

# --description--

Дано об’єкт (JavaScript) або словник (Python), визначте, чи відповідає він наступній схемі:

```json
Roles = "user" | "creator" | "moderator" | "staff" | "admin"

UserProfile = {
  username: string,
  posts: number,
  verified: boolean,
  role: Roles,
  supporter?: boolean,
  badges: string[]
}

{
  users: UserProfile[]
}
```

- Символ вертикальної риски (`|`) означає "або". role має бути одним із перелічених значень `Roles`.
- Знак питання (`?`) після supporter означає, що це поле необов’язкове, але якщо воно існує, то має бути вказаного типу.
- `UserProfile[]` позначає масив об’єктів `UserProfile`. Порожній масив є дійсним.
- Додаткові ключі дозволені

# --hints--

`isValidSchema({ users: [{ username: "ron", posts: 14, verified: true, role: "creator", badges: [ "early-adopter" ]}, { username: "cher", posts: 25, verified: true, role: "moderator", supporter: true, followers: 20, badges: [ "helper" ]}]})` має повертати `true`.

```js
assert.isTrue(isValidSchema({users: [{ username: "ron", posts: 14, verified: true, role: "creator", badges: [ "early-adopter" ]}, { username: "cher", posts: 25, verified: true, role: "moderator", supporter: true, followers: 20, badges: [ "helper" ]}]}));
```

`isValidSchema({ users: [] })` має повертати `true`.

```js
assert.isTrue(isValidSchema({ users: [] }));
```

`isValidSchema({ users: { username: "anne", posts: 0, verified: false, role: "user", supporter: false, badges: []}})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: { username: "anne", posts: 0, verified: false, role: "user", supporter: false, badges: []}}));
```

`isValidSchema({ users: [{ username: "tony", posts: 10, verified: true, role: "creator", supporter: true, badges: ["liked", 6]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "tony", posts: 10, verified: true, role: "creator", supporter: true, badges: ["liked", 6]}]}));
```

`isValidSchema({ users: [{ username: "ursula", posts: 3, verified: false, role: "user", supporter: "false", badges: ["comeback"]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "ursula", posts: 3, verified: false, role: "user", supporter: "false", badges: ["comeback"]}]}));
```

`isValidSchema({ users: [{ username: "benny", posts: 55, verified: true, role: "superstar", supporter: true, badges: ["veteran"]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "benny", posts: 55, verified: true, role: "superstar", supporter: true, badges: ["veteran"]}]}));
```

`isValidSchema({ users: [{ username: "chase", posts: 1, verified: "yes", role: "staff", supporter: false, badges: ["superstar"]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "chase", posts: 1, verified: "yes", role: "staff", supporter: false, badges: ["superstar"]}]}));
```

`isValidSchema({ users: [{ username: "carla", posts: "10", verified: false, role: "user", supporter: false, badges: ["newbie"]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "carla", posts: "10", verified: false, role: "user", supporter: false, badges: ["newbie"]}]}));
```

`isValidSchema({ users: [{ posts: 4, verified: false, role: "admin", supporter: false, badges: ["superuser", "veteran"]}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ posts: 4, verified: false, role: "admin", supporter: false, badges: ["superuser", "veteran"]}]}));
```

`isValidSchema({ users: [{ username: "harold", posts: 80, verified: true, role: "creator", supporter: true, badges: ["liked", "hero"]}, { username: "kim", posts: 11, verified: false, role: "admin", supporter: true, badges: ["first"]}, {}]})` має повертати `false`.

```js
assert.isFalse(isValidSchema({ users: [{ username: "harold", posts: 80, verified: true, role: "creator", supporter: true, badges: ["liked", "hero"]}, { username: "kim", posts: 11, verified: false, role: "admin", supporter: true, badges: ["first"]}, {}]}));
```

# --seed--

## --seed-contents--

```js
function isValidSchema(obj) {

  return obj;
}
```

# --solutions--

```js
function isValidSchema(obj) {
  const roles = ["user", "creator", "moderator", "staff", "admin"];

  function isValidUser(user) {
    return (
      typeof user.username === 'string' &&
      typeof user.posts === 'number' &&
      typeof user.verified === 'boolean' &&
      roles.includes(user.role) &&
      (user.supporter === undefined || typeof user.supporter === 'boolean') &&
      Array.isArray(user.badges) && user.badges.every(b => typeof b === 'string')
    );
  }

  return Array.isArray(obj.users) && obj.users.every(isValidUser);
}
```
