---
id: 6a15cadf5f240d05a264955b
title: "Завдання 317: калькулятор BMI"
challengeType: 28
dashedName: challenge-317
---

# --description--

За заданою вагою в фунтах і зростом в дюймах поверніть BMI (індекс маси тіла), округлений до одного десяткового знака.

Щоб отримати BMI: поділіть вагу на квадрат зросту, а потім помножте результат на 703.

# --hints--

`calculateBmi(180, 70)` має повернути `25.8`.

```js
assert.equal(calculateBmi(180, 70), 25.8);
```

`calculateBmi(140, 64)` має повернути `24.0`.

```js
assert.equal(calculateBmi(140, 64), 24.0);
```

`calculateBmi(160, 76)` має повернути `19.5`.

```js
assert.equal(calculateBmi(160, 76), 19.5);
```

`calculateBmi(200, 60)` має повернути `39.1`.

```js
assert.equal(calculateBmi(200, 60), 39.1);
```

`calculateBmi(150, 68)` має повернути `22.8`.

```js
assert.equal(calculateBmi(150, 68), 22.8);
```

# --seed--

## --seed-contents--

```js
function calculateBmi(weight, height) {

  return weight;
}
```

# --solutions--

```js
function calculateBmi(weight, height) {
  return Math.round((weight / (height * height)) * 703 * 10) / 10;
}
```
