---
id: 6a151d32d772271dd2448c2d
title: "Завдання 311: Заклинатель"
challengeType: 28
dashedName: challenge-311
---

# --description--

Дано рядок кодів заклять, які ви накладаєте. Обчисліть загальну кількість балів.

Кожен символ у рядку представляє закляття:

| Код | Закляття | Категорія | Базовий бал |
| - | - | - | - |
| `"f"` | Вогонь | Руйнування | 3 |
| `"l"` | Блискавка | Руйнування | 3 |
| `"i"` | Лід | Контроль | 2 |
| `"w"` | Вітер | Контроль | 2 |
| `"h"` | Лікування | Відновлення | 1 |
| `"s"` | Щит | Відновлення | 1 |

Комбінований множник застосовується залежно від того, скільки заклять підряд було накладено з різних категорій:

- Перше закляття завжди оцінюється за базовим значенням.
- Кожне наступне закляття з іншої категорії, ніж попереднє, збільшує множник на 1.
- Накладення закляття з тієї ж категорії, що й попереднє, скидає множник назад до 1.
- Бал за кожне закляття — це його базовий бал, помножений на поточний множник.

Поверніть загальну кількість балів за послідовність заклять.

# --hints--

`cast("fihwl")` має повернути `33`.

```js
assert.equal(cast("fihwl"), 33);
```

`cast("lwswfi")` має повернути `45`.

```js
assert.equal(cast("lwswfi"), 45);
```

`cast("wislhfl")` має повернути `37`.

```js
assert.equal(cast("wislhfl"), 37);
```

`cast("sihwlih")` має повернути `50`.

```js
assert.equal(cast("sihwlih"), 50);
```

`cast("wishlfihwslwifihl")` має повернути `101`.

```js
assert.equal(cast("wishlfihwslwifihl"), 101);
```

# --seed--

## --seed-contents--

```js
function cast(spells) {

  return spells;
}
```

# --solutions--

```js
function cast(spells) {
  const lookup = {
    f: { category: "destruction", score: 3 },
    l: { category: "destruction", score: 3 },
    i: { category: "control", score: 2 },
    w: { category: "control", score: 2 },
    h: { category: "restoration", score: 1 },
    s: { category: "restoration", score: 1 }
  };

  let total = 0;
  let multiplier = 1;
  let prevCategory = null;

  for (const spell of spells) {
    const { category, score } = lookup[spell];

    if (prevCategory && category !== prevCategory) {
      multiplier++;
    } else {
      multiplier = 1;
    }

    total += score * multiplier;
    prevCategory = category;
  }

  return total;
}
```
