---
id: 6a15cadf5f240d05a2649558
title: "Завдання 314: розклад на прості множники"
challengeType: 28
dashedName: challenge-314
---

# --description--

Дано ціле число більше за 1, поверніть його розклад на прості множники у вигляді масиву чисел у зростаючому порядку.

Розклад на прості множники — це множина простих чисел, які перемножуються, щоб отримати задане ціле число. Кожне число має саме одну таку множину. Наприклад, розклад на прості множники числа 20 — це `[2, 2, 5]`, бо 2 * 2 * 5 = 20.

Якщо задане ціле число саме є простим, поверніть його у масиві з одним елементом.

# --hints--

`primeFactorization(20)` має повернути `[2, 2, 5]`.

```js
assert.deepEqual(primeFactorization(20), [2, 2, 5]);
```

`primeFactorization(17)` має повернути `[17]`.

```js
assert.deepEqual(primeFactorization(17), [17]);
```

`primeFactorization(15)` має повернути `[3, 5]`.

```js
assert.deepEqual(primeFactorization(15), [3, 5]);
```

`primeFactorization(35)` має повернути `[5, 7]`.

```js
assert.deepEqual(primeFactorization(35), [5, 7]);
```

`primeFactorization(999)` має повернути `[3, 3, 3, 37]`.

```js
assert.deepEqual(primeFactorization(999), [3, 3, 3, 37]);
```

`primeFactorization(360)` має повернути `[2, 2, 2, 3, 3, 5]`.

```js
assert.deepEqual(primeFactorization(360), [2, 2, 2, 3, 3, 5]);
```

`primeFactorization(510510)` має повернути `[2, 3, 5, 7, 11, 13, 17]`.

```js
assert.deepEqual(primeFactorization(510510), [2, 3, 5, 7, 11, 13, 17]);
```

# --seed--

## --seed-contents--

```js
function primeFactorization(n) {

  return n;
}
```

# --solutions--

```js
function primeFactorization(n) {
  const factors = [];
  let divisor = 2;

  while (divisor * divisor <= n) {
    while (n % divisor === 0) {
      factors.push(divisor);
      n /= divisor;
    }
    divisor++;
  }

  if (n > 1) factors.push(n);

  return factors;
}
```
