---
id: 6a15cadf5f240d05a264955e
title: "Завдання 320: Банк крові"
challengeType: 28
dashedName: challenge-320
---

# --description--

Маючи масив запасів у банку крові та масив запитів пацієнтів на тип крові, поверніть рядок у форматі `"X of Y patients served"`. Де `X` — максимальна кількість пацієнтів, які можуть отримати кров із запасів банку, а `Y` — загальна кількість пацієнтів.

Кожен запис у обох масивах є одним із наступних типів крові: `"AB"`, `"A"`, `"B"` або `"O"`.

Правила сумісності:

- `"AB"` може отримувати кров від будь-якого типу.
- `"A"` може отримувати кров від `"A"` та `"O"`.
- `"B"` може отримувати кров від `"B"` та `"O"`.
- `"O"` може отримувати кров лише від `"O"`.

Повторювані записи у заданих масивах означають кількість.

# --hints--

`triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"])` має повернути `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"])` має повернути `"3 of 5 patients served"`.

```js
assert.equal(triageBlood(["A", "A", "B", "B", "AB"], ["O", "A", "B", "B", "B"]), "3 of 5 patients served");
```

`triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"])` має повернути `"4 of 5 patients served"`.

```js
assert.equal(triageBlood(["O", "A", "B", "AB"], ["AB", "AB", "AB", "AB", "AB"]), "4 of 5 patients served");
```

`triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"])` має повернути `"4 of 4 patients served"`.

```js
assert.equal(triageBlood(["O", "O", "O", "O", "O"], ["O", "A", "B", "AB"]), "4 of 4 patients served");
```

`triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"])` має повернути `"8 of 11 patients served"`.

```js
assert.equal(triageBlood(["A", "O", "B", "AB", "B", "AB", "O", "A", "A"], ["O", "A", "B", "AB", "A", "B", "A", "A", "B", "A", "B"]), "8 of 11 patients served");
```

`triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"])` має повернути `"13 of 14 patients served"`.

```js
assert.equal(triageBlood(["O", "B", "AB", "AB", "O", "A", "A", "AB", "O", "B", "B", "AB", "A", "B", "AB"], ["O", "A", "B", "B", "A", "B", "AB", "A", "B", "A", "O", "AB", "AB", "O"]), "13 of 14 patients served");
```

# --seed--

## --seed-contents--

```js
function triageBlood(bank, patients) {

  return bank;
}
```

# --solutions--

```js
function triageBlood(bank, patients) {
  const b = {};
  const p = {};
  for (const t of bank) b[t] = (b[t] || 0) + 1;
  for (const t of patients) p[t] = (p[t] || 0) + 1;

  const serve = (patient, donors) => {
    for (const donor of donors) {
      const n = Math.min(p[patient] || 0, b[donor] || 0);
      p[patient] = (p[patient] || 0) - n;
      b[donor] = (b[donor] || 0) - n;
      saved += n;
    }
  };

  let saved = 0;
  serve("O",  ["O"]);
  serve("A",  ["A", "O"]);
  serve("B",  ["B", "O"]);
  serve("AB", ["AB", "A", "B", "O"]);
  return `${saved} of ${patients.length} patients served`;
}
```
