---
id: 6a15cadf5f240d05a264955d
title: "Завдання 319: парсер frontmatter"
challengeType: 28
dashedName: challenge-319
---

# --description--

Дано рядок, що представляє блок frontmatter, проаналізуйте його та поверніть об’єкт (JavaScript) або словник (Python) з ключами та значеннями.

Frontmatter обгорнутий у роздільники `---` і містить пари `key: value`, по одній на кожній лінії. Наприклад:

```md
---
title: My Post
draft: false
views: 100
---
```

Повинно повернути:

```js
{
  title: "My Post",
  draft: false,
  views: 100
}
```

- Числа, булеві значення та рядки мають повертатися у відповідному типі.
- Даний рядок матиме нові лінії, розділені символом нового рядка (`"\n"`). Наведений вище приклад буде поданий як: `"---\ntitle: My Post\ndraft: false\nviews: 100\n---"`.

# --hints--

`parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---")` має повернути `{ title: "My Post", draft: false, views: 100 }`.

```js
assert.deepEqual(parseFrontmatter("---\ntitle: My Post\ndraft: false\nviews: 100\n---"), { title: "My Post", draft: false, views: 100 });
```

`parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---")` має повернути `{ id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false }`.

```js
assert.deepEqual(parseFrontmatter("---\nid: 6a174db57256a112f932195c\ntitle: My Book\nlocale: en\nwordCount: 10000\npublished: false\n---"), { id: "6a174db57256a112f932195c", title: "My Book", locale: "en", wordCount: 10000, published: false });
```

`parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---")` має повернути `{ version: "1.0.0", url: "https://example.com", private: true }`.

```js
assert.deepEqual(parseFrontmatter("---\nversion: 1.0.0\nurl: https://example.com\nprivate: true\n---"), { version: "1.0.0", url: "https://example.com", private: true });
```

`parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---")` має повернути `{ rating: 4.5, price: 9.99 }`.

```js
assert.deepEqual(parseFrontmatter("---\nrating: 4.5\nprice: 9.99\n---"), { rating: 4.5, price: 9.99 });
```

# --seed--

## --seed-contents--

```js
function parseFrontmatter(str) {

  return str;
}
```

# --solutions--

```js
function parseFrontmatter(str) {
  const lines = str.split("\n").filter(line => line !== "---" && line !== "");

  return lines.reduce((obj, line) => {
    const [key, ...rest] = line.split(":");
    const value = rest.join(":").trim();

    if (value === "true") obj[key.trim()] = true;
    else if (value === "false") obj[key.trim()] = false;
    else if (!isNaN(value)) obj[key.trim()] = Number(value);
    else obj[key.trim()] = value;

    return obj;
  }, {});
}
```
