---
id: 6a1d9f98e819ed70a0e994e0
title: "Завдання 334: Точна здача"
challengeType: 28
dashedName: challenge-334
---

# --description--

Дано ціле число в центах, поверніть кількість різних способів дати точну здачу, використовуючи пенні (1 цент), нікелі (5 центів), дайми (10 центів) і квартери (25 центів).

# --hints--

`exactChange(3)` має повернути `1`.

```js
assert.equal(exactChange(3), 1);
```

`exactChange(9)` має повернути `2`.

```js
assert.equal(exactChange(9), 2);
```

`exactChange(17)` має повернути `6`.

```js
assert.equal(exactChange(17), 6);
```

`exactChange(39)` має повернути `24`.

```js
assert.equal(exactChange(39), 24);
```

`exactChange(61)` має повернути `73`.

```js
assert.equal(exactChange(61), 73);
```

`exactChange(99)` має повернути `213`.

```js
assert.equal(exactChange(99), 213);
```

# --seed--

## --seed-contents--

```js
function exactChange(amount) {

  return amount;
}
```

# --solutions--

```js
function exactChange(amount) {
  const coins = [1, 5, 10, 25];
  const dp = new Array(amount + 1).fill(0);
  dp[0] = 1;

  for (const coin of coins) {
    for (let i = coin; i <= amount; i++) {
      dp[i] += dp[i - coin];
    }
  }

  return dp[amount];
}
```
