---
id: 6a1d9f98e819ed70a0e994dc
title: "Завдання 330: слова в нижньому регістрі"
challengeType: 28
dashedName: challenge-330
---

# --description--

Дано рядок, поверніть лише ті слова, які повністю в нижньому регістрі, у їхньому початковому порядку та з пробілом між кожним словом.

# --hints--

`getLowercaseWords("hello GOOD world")` має повернути `"hello world"`.

```js
assert.equal(getLowercaseWords("hello GOOD world"), "hello world");
```

`getLowercaseWords("these are all lowercase")` має повернути `"these are all lowercase"`.

```js
assert.equal(getLowercaseWords("these are all lowercase"), "these are all lowercase");
```

`getLowercaseWords("less is NoT more")` має повернути `"less is more"`.

```js
assert.equal(getLowercaseWords("less is NoT more"), "less is more");
```

`getLowercaseWords("DonT eat pizza every OTHER day")` має повернути `"eat pizza every day"`.

```js
assert.equal(getLowercaseWords("DonT eat pizza every OTHER day"), "eat pizza every day");
```

`getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog")` має повернути `"the quick brown fox jumped over the lazy dog"`.

```js
assert.equal(getLowercaseWords("the Super quick AND snEaky brown fox Leapt anD jumped over aNd AROUND the lazy SloW dog"), "the quick brown fox jumped over the lazy dog");
```

# --seed--

## --seed-contents--

```js
function getLowercaseWords(str) {

  return str;
}
```

# --solutions--

```js
function getLowercaseWords(str) {
  return str.split(" ").filter(word => word === word.toLowerCase()).join(" ");
}
```
