---
id: 6a1d9f98e819ed70a0e994da
title: "Завдання 328: алгоритм Капрекара"
challengeType: 28
dashedName: challenge-328
---

# --description--

Дано 4-цифрове число, поверніть кількість разів, які потрібно застосувати алгоритм Капрекара, щоб дійти до 6174.

Алгоритм Капрекара працює так:

- Впорядкуйте цифри за спаданням, щоб утворити найбільше число
- Впорядкуйте цифри за зростанням, щоб утворити найменше число (за потреби додайте провідні нулі)
- Відніміть менше число від більшого
- Повторюйте з новим числом

# --hints--

`kaprekar(1234)` має повернути `3`.

```js
assert.equal(kaprekar(1234), 3);
```

`kaprekar(2025)` має повернути `6`.

```js
assert.equal(kaprekar(2025), 6);
```

`kaprekar(7173)` має повернути `4`.

```js
assert.equal(kaprekar(7173), 4);
```

`kaprekar(3164)` має повернути `7`.

```js
assert.equal(kaprekar(3164), 7);
```

`kaprekar(8082)` має повернути `2`.

```js
assert.equal(kaprekar(8082), 2);
```

# --seed--

## --seed-contents--

```js
function kaprekar(n) {

  return n;
}
```

# --solutions--

```js
function kaprekar(n) {
  let steps = 0;
  let current = n;

  while (current !== 6174) {
    const digits = String(current).padStart(4, "0").split("").map(Number);
    const desc = parseInt(digits.slice().sort((a, b) => b - a).join(""));
    const asc = parseInt(digits.slice().sort((a, b) => a - b).join(""));
    current = desc - asc;
    steps++;
  }

  return steps;
}
```
