---
id: 6a1d9f98e819ed70a0e994dd
title: "Завдання 331: найближчий кратний"
challengeType: 28
dashedName: challenge-331
---

# --description--

Дано два цілі числа, округліть перше до найближчого кратного другого.

# --hints--

`roundToNearestMultiple(5, 3)` має повернути `6`.

```js
assert.equal(roundToNearestMultiple(5, 3), 6);
```

`roundToNearestMultiple(17, 4)` має повернути `16`.

```js
assert.equal(roundToNearestMultiple(17, 4), 16);
```

`roundToNearestMultiple(43, 5)` має повернути `45`.

```js
assert.equal(roundToNearestMultiple(43, 5), 45);
```

`roundToNearestMultiple(38, 11)` має повернути `33`.

```js
assert.equal(roundToNearestMultiple(38, 11), 33);
```

`roundToNearestMultiple(93, 12)` має повернути `96`.

```js
assert.equal(roundToNearestMultiple(93, 12), 96);
```

# --seed--

## --seed-contents--

```js
function roundToNearestMultiple(num, multiple) {

  return num;
}
```

# --solutions--

```js
function roundToNearestMultiple(num, multiple) {
  return Math.round(num / multiple) * multiple;
}
```
