---
id: 6a1d9f98e819ed70a0e994de
title: "Завдання 332: сортування проблем"
challengeType: 28
dashedName: challenge-332
---

# --description--

Враховуючи кількість мілісекунд від останнього повідомлення у проблемі та текст останнього повідомлення, визначте, що слід зробити з проблемою за цими правилами:

- Якщо останнє повідомлення було менше ніж 7 днів тому, поверніть `"leave it"`
- Якщо останнє повідомлення було 7 або більше днів тому і його вміст містить `"bump"` (без урахування регістру), поверніть `"close it"`
- Інакше поверніть `"bump it"`

# --hints--

`triageIssue(86400000, "Lets fix it")` має повернути `"leave it"`.

```js
assert.equal(triageIssue(86400000, "Lets fix it"), "leave it");
```

`triageIssue(1209600000, "still waiting")` має повернути `"bump it"`.

```js
assert.equal(triageIssue(1209600000, "still waiting"), "bump it");
```

`triageIssue(864000000, "bump")` має повернути `"close it"`.

```js
assert.equal(triageIssue(864000000,  "bump"), "close it");
```

`triageIssue(604800000, "Do we still want this?")` має повернути `"bump it"`.

```js
assert.equal(triageIssue(604800000, "Do we still want this?"), "bump it");
```

`triageIssue(604800000, "Bumping this")` має повернути `"close it"`.

```js
assert.equal(triageIssue(604800000, "Bumping this"), "close it");
```

`triageIssue(345600000, "I'll make a PR")` має повернути `"leave it"`.

```js
assert.equal(triageIssue(345600000, "I'll make a PR"), "leave it");
```

# --seed--

## --seed-contents--

```js
function triageIssue(ms, message) {

  return ms;
}
```

# --solutions--

```js
function triageIssue(ms, message) {
  const sevenDays = 7 * 24 * 60 * 60 * 1000;

  if (ms < sevenDays) return "leave it";
  if (message.toLowerCase().includes("bump")) return "close it";
  return "bump it";
}
```
