---
id: 6a1d9f98e819ed70a0e994e1
title: "Завдання 335: П’ять кубиків"
challengeType: 28
dashedName: challenge-335
---

# --description--

Дано масив із п’яти кубиків зі значеннями від 1 до 6, поверніть найкращу можливу комбінацію.

Ось комбінації, впорядковані від найнижчої до найвищої:

| Комбінація | Опис |
|------|-------------|
| `"no pair"` | Немає пари або кращої комбінації |
| `"pair"` | Два кубики з однаковим значенням |
| `"two pair"` | Дві різні пари |
| `"three of a kind"` | Три кубики з однаковим значенням |
| `"small straight"` | Чотири послідовні значення |
| `"large straight"` | П’ять послідовних значень |
| `"full house"` | Три однакові та одна пара |
| `"four of a kind"` | Чотири кубики з однаковим значенням |
| `"five of a kind"` | Всі п’ять кубиків з однаковим значенням |

# --hints--

`fiveDice([1, 1, 1, 1, 1])` має повернути `"five of a kind"`.

```js
assert.equal(fiveDice([1, 1, 1, 1, 1]), "five of a kind");
```

`fiveDice([5, 5, 5, 6, 5])` має повернути `"four of a kind"`.

```js
assert.equal(fiveDice([5, 5, 5, 6, 5]), "four of a kind");
```

`fiveDice([2, 5, 6, 4, 3])` має повернути `"large straight"`.

```js
assert.equal(fiveDice([2, 5, 6, 4, 3]), "large straight");
```

`fiveDice([4, 3, 3, 3, 1])` має повернути `"three of a kind"`.

```js
assert.equal(fiveDice([4, 3, 3, 3, 1]), "three of a kind");
```

`fiveDice([4, 6, 2, 6, 5])` має повернути `"pair"`.

```js
assert.equal(fiveDice([4, 6, 2, 6, 5]), "pair");
```

`fiveDice([1, 4, 5, 6, 2])` має повернути `"no pair"`.

```js
assert.equal(fiveDice([1, 4, 5, 6, 2]), "no pair");
```

`fiveDice([1, 3, 4, 6, 2])` має повернути `"small straight"`.

```js
assert.equal(fiveDice([1, 3, 4, 6, 2]), "small straight");
```

`fiveDice([2, 2, 5, 2, 5])` має повернути `"full house"`.

```js
assert.equal(fiveDice([2, 2, 5, 2, 5]), "full house");
```

`fiveDice([6, 4, 5, 6, 4])` має повернути `"two pair"`.

```js
assert.equal(fiveDice([6, 4, 5, 6, 4]), "two pair");
```

# --seed--

## --seed-contents--

```js
function fiveDice(dice) {

  return dice;
}
```

# --solutions--

```js
function fiveDice(dice) {
  const counts = {};
  for (const d of dice) counts[d] = (counts[d] || 0) + 1;
  const values = Object.values(counts).sort((a, b) => b - a);
  const keys = Object.keys(counts).map(Number).sort((a, b) => a - b);

  const isLargeStraight = keys.length === 5 && keys[4] - keys[0] === 4;
  const isSmallStraight = !isLargeStraight && (() => {
    for (let i = 0; i <= 1; i++) {
      const slice = keys.slice(i, i + 4);
      if (slice.length === 4 && slice[3] - slice[0] === 3 &&
        new Set(slice).size === 4) return true;
    }
    return false;
  })();

  if (values[0] === 5) return "five of a kind";
  if (values[0] === 4) return "four of a kind";
  if (values[0] === 3 && values[1] === 2) return "full house";
  if (isLargeStraight) return "large straight";
  if (isSmallStraight) return "small straight";
  if (values[0] === 3) return "three of a kind";
  if (values[0] === 2 && values[1] === 2) return "two pair";
  if (values[0] === 2) return "pair";
  return "no pair";
}
```
