---
id: 6a2037a68a0bc2aef0006007
title: "Завдання 345: Змішувач слів"
challengeType: 28
dashedName: challenge-345
---

# --description--

Дано два слова, поверніть нове слово, поєднавши першу половину першого слова з другою половиною другого слова.

- Для слів непарної довжини перша половина є коротшою половиною.

# --hints--

`blendWords("turtle", "toucan")` має повернути `"turcan"`.

```js
assert.equal(blendWords("turtle", "toucan"), "turcan");
```

`blendWords("chipmunk", "flamingo")` має повернути `"chipingo"`.

```js
assert.equal(blendWords("chipmunk", "flamingo"), "chipingo");
```

`blendWords("falcon", "pelican")` має повернути `"falican"`.

```js
assert.equal(blendWords("falcon", "pelican"), "falican");
```

`blendWords("hyena", "iguana")` має повернути `"hyana"`.

```js
assert.equal(blendWords("hyena", "iguana"), "hyana");
```

`blendWords("scorpion", "gorilla")` має повернути `"scorilla"`.

```js
assert.equal(blendWords("scorpion", "gorilla"), "scorilla");
```

`blendWords("platypus", "wolverine")` має повернути `"platerine"`.

```js
assert.equal(blendWords("platypus", "wolverine"), "platerine");
```

# --seed--

## --seed-contents--

```js
function blendWords(word1, word2) {

  return word1;
}
```

# --solutions--

```js
function blendWords(word1, word2) {
  const firstHalf = word1.slice(0, Math.floor(word1.length / 2));
  const secondHalf = word2.slice(Math.floor(word2.length / 2));
  return firstHalf + secondHalf;
}
```
