---
id: 6a2037a68a0bc2aef0006002
title: "Завдання 340: конвертер Pig Latin"
challengeType: 28
dashedName: challenge-340
---

# --description--

Дано рядок, перетворіть його на Pig Latin за такими правилами:

- Якщо слово починається з голосної (`"a"`, `"e"`, `"i"`, `"o"` або `"u"`), додайте в кінці `"way"`. Наприклад, `"universe"` перетворюється на `"universeway"`.
- Якщо слово починається з однієї або кількох приголосних, перенесіть їх у кінець і додайте `"ay"`. Наприклад, `"hello"` перетворюється на `"ellohay"`.
- Збережіть регістр першої літери. Наприклад, `"Hello"` перетворюється на `"Ellohay"`.

# --hints--

`pigLatin("universe")` має повернути `"universeway"`.

```js
assert.equal(pigLatin("universe"), "universeway");
```

`pigLatin("hello")` має повернути `"ellohay"`.

```js
assert.equal(pigLatin("hello"), "ellohay");
```

`pigLatin("hello universe")` має повернути `"ellohay universeway"`.

```js
assert.equal(pigLatin("hello universe"), "ellohay universeway");
```

`pigLatin("Hello universe")` має повернути `"Ellohay universeway"`.

```js
assert.equal(pigLatin("Hello universe"), "Ellohay universeway");
```

`pigLatin("Pig Latin is fun")` має повернути `"Igpay Atinlay isway unfay"`.

```js
assert.equal(pigLatin("Pig Latin is fun"), "Igpay Atinlay isway unfay");
```

`pigLatin("The quick brown fox jumped over the lazy dog")` має повернути `"Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday"`.

```js
assert.equal(pigLatin("The quick brown fox jumped over the lazy dog"), "Ethay uickqay ownbray oxfay umpedjay overway ethay azylay ogday");
```

# --seed--

## --seed-contents--

```js
function pigLatin(str) {

  return str;
}
```

# --solutions--

```js
function pigLatin(str) {
  const vowels = "aeiou";

  return str.split(" ").map(word => {
    const lowerWord = word.toLowerCase();

    if (vowels.includes(lowerWord[0])) {
      return word + "way";
    }

    let i = 0;
    while (i < word.length && !vowels.includes(lowerWord[i])) {
      i++;
    }

    const consonants = word.slice(0, i);
    const rest = word.slice(i);
    const result = rest + consonants + "ay";

    if (word[0] === word[0].toUpperCase()) {
      return result[0].toUpperCase() + result.slice(1).toLowerCase();
    }

    return result;
  }).join(" ");
}
```
