---
id: 6a2037a68a0bc2aef0006006
title: "Завдання 344: золотий перетин"
challengeType: 28
dashedName: challenge-344
---

# --description--

Дано два числа, визначте, чи їхнє співвідношення приблизно дорівнює золотому перетину.

- Використовуйте золотий перетин `1.618`
- Дозвольте похибку `0.01`

# --hints--

`isGoldenRatio(21, 34)` має повернути `true`.

```js
assert.isTrue(isGoldenRatio(21, 34));
```

`isGoldenRatio(15, 20)` має повернути `false`.

```js
assert.isFalse(isGoldenRatio(15, 20));
```

`isGoldenRatio(8, 13)` має повернути `true`.

```js
assert.isTrue(isGoldenRatio(8, 13));
```

`isGoldenRatio(10, 16)` має повернути `false`.

```js
assert.isFalse(isGoldenRatio(10, 16));
```

`isGoldenRatio(1618, 1000)` має повернути `true`.

```js
assert.isTrue(isGoldenRatio(1618, 1000));
```

`isGoldenRatio(88, 55)` має повернути `false`.

```js
assert.isFalse(isGoldenRatio(88, 55));
```

# --seed--

## --seed-contents--

```js
function isGoldenRatio(a, b) {

  return a;
}
```

# --solutions--

```js
function isGoldenRatio(a, b) {
  const goldenRatio = 1.618;
  const ratio = Math.max(a, b) / Math.min(a, b);
  return Math.abs(ratio - goldenRatio) <= 0.01;
}
```
