---
id: 6a2037a68a0bc2aef0005fff
title: "Завдання 337: фішка"
challengeType: 28
dashedName: challenge-337
---

# --description--

Дано рядок з фішками, поверніть загальну кількість, яку вони представляють.

- Кожна вертикальна риска `"|"` представляє одну фішку.
- Кожна п’ята позначка представлена як коса риска `"/"`, що завершує групу з п’яти (`"||||/"`).
- Групи розділені пробілом.

# --hints--

`getTallyCount("||||")` має повернути `4`.

```js
assert.equal(getTallyCount("||||"), 4);
```

`getTallyCount("||||/")` має повернути `5`.

```js
assert.equal(getTallyCount("||||/"), 5);
```

`getTallyCount("||||/ |||")` має повернути `8`.

```js
assert.equal(getTallyCount("||||/ |||"), 8);
```

`getTallyCount("||||/ ||||/ ||||/ ||")` має повернути `17`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||"), 17);
```

`getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |")` має повернути `41`.

```js
assert.equal(getTallyCount("||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ ||||/ |"), 41);
```

# --seed--

## --seed-contents--

```js
function getTallyCount(str) {

  return str;
}
```

# --solutions--

```js
function getTallyCount(str) {
  return str.replace(/ /g, "").length;
}
```
