---
id: 6a2037a68a0bc2aef0006000
title: "Завдання 338: калькулятор віку домашніх тварин"
challengeType: 28
dashedName: challenge-338
---

# --description--

Враховуючи тип домашньої тварини та її вік у людських роках, поверніть еквівалентний вік у роках тварини, використовуючи таку таблицю конвертації:

| Домашня тварина | Множник |
|-----|-----------|
| `"dog"` | 7 |
| `"cat"` | 6 |
| `"rabbit"` | 8 |
| `"hamster"` | 30 |
| `"guinea pig"` | 12 |
| `"goldfish"` | 6 |
| `"bird"` | 5 |

# --hints--

`petYears("dog", 5)` має повернути `35`.

```js
assert.equal(petYears("dog", 5), 35);
```

`petYears("cat", 9)` має повернути `54`.

```js
assert.equal(petYears("cat", 9), 54);
```

`petYears("rabbit", 3)` має повернути `24`.

```js
assert.equal(petYears("rabbit", 3), 24);
```

`petYears("hamster", 4)` має повернути `120`.

```js
assert.equal(petYears("hamster", 4), 120);
```

`petYears("guinea pig", 5)` має повернути `60`.

```js
assert.equal(petYears("guinea pig", 5), 60);
```

`petYears("goldfish", 2)` має повернути `12`.

```js
assert.equal(petYears("goldfish", 2), 12);
```

`petYears("bird", 1)` має повернути `5`.

```js
assert.equal(petYears("bird", 1), 5);
```

# --seed--

## --seed-contents--

```js
function petYears(pet, age) {

  return pet;
}
```

# --solutions--

```js
function petYears(pet, age) {
  const table = {
    dog: 7,
    cat: 6,
    rabbit: 8,
    hamster: 30,
    "guinea pig": 12,
    goldfish: 6,
    bird: 5,
  };

  return age * table[pet];
}
```
