---
id: 6a2037a68a0bc2aef0006008
title: "Завдання 346: скарбничка"
challengeType: 28
dashedName: challenge-346
---

# --description--

Дано об’єкт, який представляє скарбничку, поверніть загальне значення у вигляді рядка, відформатованого як `"$D.CC"`.

Об’єкт може містити будь-які з наступних:

| Монета | Значення |
|------|-------|
| `pennies` | $0.01 |
| `nickels` | $0.05 |
| `dimes` | $0.10 |
| `quarters` | $0.25 |

# --hints--

`piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 })` має повернути `"$1.98"`.

```js
assert.equal(piggyBank({ pennies: 3, nickels: 5, dimes: 2, quarters: 6 }), "$1.98");
```

`piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 })` має повернути `"$0.41"`.

```js
assert.equal(piggyBank({ pennies: 1, nickels: 1, dimes: 1, quarters: 1 }), "$0.41");
```

`piggyBank({ nickels: 8, dimes: 6, quarters: 5 })` має повернути `"$2.25"`.

```js
assert.equal(piggyBank({ nickels: 8, dimes: 6, quarters: 5 }), "$2.25");
```

`piggyBank({})` має повернути `"$0.00"`.

```js
assert.equal(piggyBank({}), "$0.00");
```

`piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 })` має повернути `"$6.76"`.

```js
assert.equal(piggyBank({ pennies: 146, nickels: 11, dimes: 0, quarters: 19 }), "$6.76");
```

# --seed--

## --seed-contents--

```js
function piggyBank(coins) {

  return coins;
}
```

# --solutions--

```js
function piggyBank(coins) {
  const values = {
    pennies: 0.01,
    nickels: 0.05,
    dimes: 0.10,
    quarters: 0.25,
  };

  const total = Object.entries(coins).reduce((sum, [coin, count]) => {
    return sum + (values[coin] * count);
  }, 0);

  return "$" + total.toFixed(2);
}
```
