---
id: 6a22d77ddf034bc4e35b1d57
title: "Завдання 350: Відстань між літерами"
challengeType: 28
dashedName: challenge-350
---

# --description--

Дано два рядки однакової довжини, поверніть суму найкоротших відстаней між кожною парою символів.

- Вхідні дані міститимуть лише літери нижнього регістру
- Алфавіт розглядається як коло, тому відстань між `a` та `z` дорівнює 1.

# --hints--

`letterDistance("abc", "bcd")` має повернути `3`.

```js
assert.equal(letterDistance("abc", "bcd"), 3);
```

`letterDistance("abc", "xyz")` має повернути `9`.

```js
assert.equal(letterDistance("abc", "xyz"), 9);
```

`letterDistance("encrypt", "decrypt")` має повернути `10`.

```js
assert.equal(letterDistance("encrypt", "decrypt"), 10);
```

`letterDistance("algorithm", "codeblock")` має повернути `43`.

```js
assert.equal(letterDistance("algorithm", "codeblock"), 43);
```

`letterDistance("lobster", "penguin")` має повернути `47`.

```js
assert.equal(letterDistance("lobster", "penguin"), 47);
```

`letterDistance("alligator", "crocodile")` має повернути `55`.

```js
assert.equal(letterDistance("alligator", "crocodile"), 55);
```

# --seed--

## --seed-contents--

```js
function letterDistance(str1, str2) {

  return str1;
}
```

# --solutions--

```js
function letterDistance(str1, str2) {
  return str1.split("").reduce((sum, char, i) => {
    const diff = Math.abs(char.charCodeAt(0) - str2[i].charCodeAt(0));
    return sum + Math.min(diff, 26 - diff);
  }, 0);
}
```
