---
id: 6a22d77ddf034bc4e35b1d54
title: "Завдання 347: теорія ігор"
challengeType: 28
dashedName: challenge-347
---

# --description--

Дано два рядки однакової довжини, які представляють стратегії двох гравців у грі, поверніть рахунки у вигляді масиву `[player1, player2]`.

- Дано рядки міститимуть лише одну з двох літер: `"C"` (співпрацювати) або `"D"` (зраджувати).
- Кожен символ представляє один раунд, оцінений так:
  - Якщо обидва гравці співпрацюють, кожен отримує 3 очки.
  - Якщо обидва гравці зраджують, кожен отримує 1 очко.
  - Якщо один гравець зраджує, а інший співпрацює, зрадник отримує 5 очок, а співпрацюючий — 0.

# --hints--

`playGame("CCCC", "CCCC")` має повернути `[12, 12]`.

```js
assert.deepEqual(playGame("CCCC", "CCCC"), [12, 12]);
```

`playGame("DDDD", "DDDD")` має повернути `[4, 4]`.

```js
assert.deepEqual(playGame("DDDD", "DDDD"), [4, 4]);
```

`playGame("CCDD", "CDDD")` має повернути `[5, 10]`.

```js
assert.deepEqual(playGame("CCDD", "CDDD"), [5, 10]);
```

`playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD")` має повернути `[24, 34]`.

```js
assert.deepEqual(playGame("CCCDCDCCCDDC", "CCDDCDCDDCCD"), [24, 34]);
```

`playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC")` має повернути `[66, 21]`.

```js
assert.deepEqual(playGame("DDCCDDDDCDDCDDDCDD", "CCDCCCDCCCDCCCCDCC"), [66, 21]);
```

# --seed--

## --seed-contents--

```js
function playGame(p1, p2) {

  return p1;
}
```

# --solutions--

```js
function playGame(p1, p2) {
  let scores = [0, 0];

  for (let i = 0; i < p1.length; i++) {
    if (p1[i] === 'C' && p2[i] === 'C') {
      scores[0] += 3; scores[1] += 3;
    } else if (p1[i] === 'D' && p2[i] === 'D') {
      scores[0] += 1; scores[1] += 1;
    } else if (p1[i] === 'D' && p2[i] === 'C') {
      scores[0] += 5; scores[1] += 0;
    } else {
      scores[0] += 0; scores[1] += 5;
    }
  }

  return scores;
}
```
