---
id: 6a22d77ddf034bc4e35b1d5b
title: "Завдання 354: рейтинг контрастності 3"
challengeType: 28
dashedName: challenge-354
---

# --description--

Дано два масиви, які представляють значення RGB, і булевий параметр, що вказує, чи є текст великим, поверніть рейтинг контрастності WCAG, використовуючи наступний метод:

Спочатку перетворіть кожне значення RGB у відносну яскравість:

- Поділіть кожен канал `[R, G, B]` на 255, щоб отримати значення від `0` до `1`
- Застосуйте формулу гамма-корекції до кожного каналу:
  - Якщо значення каналу менше або дорівнює `0.04045`: `channel / 12.92`
  - Інакше: `((channel + 0.055) / 1.055) ^ 2.4`
- Обчисліть яскравість: `0.2126 * R + 0.7152 * G + 0.0722 * B`

Потім обчисліть коефіцієнт контрастності, додавши 0.05 до кожного значення яскравості, а потім поділивши світліше на темніше. Світліше завжди буде першим аргументом.

Поверніть рейтинг на основі коефіцієнта контрастності, використовуючи таку таблицю:

| Рейтинг | Звичайний текст | Великий текст |
|--------|-----------------|---------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | нижче 4.5 | нижче 3.0 |

# --hints--

`getContrastRating([255, 255, 255], [0, 0, 0], false)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating([255, 255, 255], [0, 0, 0], false), "AAA");
```

`getContrastRating([215, 188, 188], [55, 55, 55], false)` має повернути `"AA"`.

```js
assert.equal(getContrastRating([215, 188, 188], [55, 55, 55], false), "AA");
```

`getContrastRating([143, 144, 210], [46, 47, 61], false)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating([143, 144, 210], [46, 47, 61], false), "Fail");
```

`getContrastRating([167, 167, 210], [53, 10, 53], true)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating([167, 167, 210], [53, 10, 53], true), "AAA");
```

`getContrastRating([135, 147, 155], [60, 70, 90], true)` має повернути `"AA"`.

```js
assert.equal(getContrastRating([135, 147, 155], [60, 70, 90], true), "AA");
```

`getContrastRating([125, 210, 195], [105, 130, 90], true)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating([125, 210, 195], [105, 130, 90], true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {

  return rgb1;
}
```

# --solutions--

```js
function getContrastRating(rgb1, rgb2, isLargeText) {
  function toLuminance([r, g, b]) {
    return [r, g, b].map(channel => {
      channel = channel / 255;
      return channel <= 0.04045
        ? channel / 12.92
        : ((channel + 0.055) / 1.055) ** 2.4;
    }).reduce((sum, c, i) => sum + c * [0.2126, 0.7152, 0.0722][i], 0);
  }

  const l1 = toLuminance(rgb1);
  const l2 = toLuminance(rgb2);
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
