---
id: 6a22d77ddf034bc4e35b1d56
title: "Завдання 349: Сигнал стільникового зв’язку"
challengeType: 28
dashedName: challenge-349
---

# --description--

Дано сітку з трьома показниками стільникових веж, визначте місцезнаходження телефону.

- Кожна клітинка в сітці містить або `0` (немає вежі), або додатне ціле число, яке представляє кількість клітинок до телефону, виміряну по прямій лінії: горизонтальній, вертикальній або діагональній.
- Поверніть `[row, col]` клітинки, яка знаходиться на правильній відстані у клітинках від усіх трьох веж.
- Завжди існує рівно одне рішення.

# --hints--

`findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]])` має повернути `[1, 2]`.

```js
assert.deepEqual(findSignal([[0, 0, 1], [0, 1, 0], [0, 0, 1]]), [1, 2]);
```

`findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]])` має повернути `[2, 1]`.

```js
assert.deepEqual(findSignal([[0, 2, 0], [1, 0, 0], [0, 0, 1]]), [2, 1]);
```

`findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]])` має повернути `[2, 2]`.

```js
assert.deepEqual(findSignal([[0, 0, 2, 0], [0, 0, 0, 0], [2, 0, 0, 0], [0, 0, 0, 1]]), [2, 2]);
```

`findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]])` має повернути `[3, 4]`.

```js
assert.deepEqual(findSignal([[0, 3, 0, 0, 0], [0, 0, 0, 0, 2], [0, 0, 0, 0, 0], [4, 0, 0, 0, 0], [0, 0, 0, 0, 0]]), [3, 4]);
```

`findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]])` має повернути `[3, 3]`.

```js
assert.deepEqual(findSignal([[3, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 0, 0, 0, 0, 0], [0, 2, 0, 0, 0, 2]]), [3, 3]);
```

# --seed--

## --seed-contents--

```js
function findSignal(grid) {

  return grid;
}
```

# --solutions--

```js
function findSignal(grid) {
  const rows = grid.length;
  const cols = grid[0].length;

  const towers = [];
  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) towers.push([r, c, grid[r][c]]);
    }
  }

  function isValid(r, c, tr, tc, dist) {
    const dr = Math.abs(r - tr);
    const dc = Math.abs(c - tc);
    return (dr === dist || dc === dist) && (dr === 0 || dc === 0 || dr === dc);
  }

  for (let r = 0; r < rows; r++) {
    for (let c = 0; c < cols; c++) {
      if (grid[r][c] !== 0) continue;
      if (towers.every(([tr, tc, dist]) => isValid(r, c, tr, tc, dist))) {
        return [r, c];
      }
    }
  }
}
```
