---
id: 6a22d77ddf034bc4e35b1d5a
title: "Завдання 353: рейтинг контрастності 2"
challengeType: 28
dashedName: challenge-353
---

# --description--

Маючи два значення відносної яскравості та булевий параметр, що вказує, чи є текст великим, поверніть рейтинг контрастності WCAG, використовуючи наступний метод:

Обчисліть коефіцієнт контрастності, додавши 0.05 до кожного значення яскравості, а потім поділивши світліше значення на темніше. Світліше значення завжди буде першим аргументом.

Поверніть рейтинг на основі коефіцієнта контрастності, використовуючи таку таблицю:

| Рейтинг | Звичайний текст | Великий текст |
|--------|-----------------|---------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | нижче 4.5 | нижче 3.0 |

# --hints--

`getContrastRating(1.0, 0.0, false)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating(1.0, 0.0, false), "AAA");
```

`getContrastRating(0.9015, 0.1364, false)` має повернути `"AA"`.

```js
assert.equal(getContrastRating(0.9015, 0.1364, false), "AA");
```

`getContrastRating(0.8965, 0.1628, false)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating(0.8965, 0.1628, false), "Fail");
```

`getContrastRating(0.7469, 0.0957, true)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating(0.7469, 0.0957, true), "AAA");
```

`getContrastRating(0.7489, 0.2018, true)` має повернути `"AA"`.

```js
assert.equal(getContrastRating(0.7489, 0.2018, true), "AA");
```

`getContrastRating(0.6571, 0.1974, true)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating(0.6571, 0.1974, true), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(l1, l2, isLargeText) {

  return l1;
}
```

# --solutions--

```js
function getContrastRating(l1, l2, isLargeText) {
  const ratio = (l1 + 0.05) / (l2 + 0.05);

  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
