---
id: 6a22d77ddf034bc4e35b1d59
title: "Завдання 352: рейтинг коефіцієнта контрастності 1"
challengeType: 28
dashedName: challenge-352
---

# --description--

За заданим коефіцієнтом контрастності та булевим значенням, яке вказує, чи є текст великим, поверніть рейтинг WCAG, використовуючи таку таблицю:

| Рейтинг | Звичайний текст | Великий текст |
|--------|-----------------|---------------|
| `"AAA"` | 7.0+ | 4.5+ |
| `"AA"` | 4.5+ | 3.0+ |
| `"Fail"` | нижче 4.5 | нижче 3.0 |

# --hints--

`getContrastRating("7.5", false)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating("7.5", false), "AAA");
```

`getContrastRating("4.8", false)` має повернути `"AA"`.

```js
assert.equal(getContrastRating("4.8", false), "AA");
```

`getContrastRating("4.2", false)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating("4.2", false), "Fail");
```

`getContrastRating("4.5", true)` має повернути `"AAA"`.

```js
assert.equal(getContrastRating("4.5", true), "AAA");
```

`getContrastRating("3.0", true)` має повернути `"AA"`.

```js
assert.equal(getContrastRating("3.0", true), "AA");
```

`getContrastRating("2.7", false)` має повернути `"Fail"`.

```js
assert.equal(getContrastRating("2.7", false), "Fail");
```

# --seed--

## --seed-contents--

```js
function getContrastRating(ratio, isLargeText) {

  return ratio;
}
```

# --solutions--

```js
function getContrastRating(ratio, isLargeText) {
  if (isLargeText) {
    if (ratio >= 4.5) return "AAA";
    if (ratio >= 3.0) return "AA";
  } else {
    if (ratio >= 7.0) return "AAA";
    if (ratio >= 4.5) return "AA";
  }
  return "Fail";
}
```
