---
id: 6a22d77ddf034bc4e35b1d58
title: "Завдання 351: пронічне число"
challengeType: 28
dashedName: challenge-351
---

# --description--

Дано число, визначте, чи є воно пронічним числом.

Пронічне число — це добуток двох послідовних цілих чисел. Наприклад, 6 є пронічним, бо 2 * 3 = 6.

# --hints--

`isPronic(6)` має повернути `true`.

```js
assert.isTrue(isPronic(6));
```

`isPronic(15)` має повернути `false`.

```js
assert.isFalse(isPronic(15));
```

`isPronic(12)` має повернути `true`.

```js
assert.isTrue(isPronic(12));
```

`isPronic(132)` має повернути `true`.

```js
assert.isTrue(isPronic(132));
```

`isPronic(80)` має повернути `false`.

```js
assert.isFalse(isPronic(80));
```

`isPronic(0)` має повернути `true`.

```js
assert.isTrue(isPronic(0));
```

# --seed--

## --seed-contents--

```js
function isPronic(n) {

  return n;
}
```

# --solutions--

```js
function isPronic(n) {
  const k = Math.floor(Math.sqrt(n));
  return k * (k + 1) === n;
}
```
