---
id: 6a26df95efa55a2524399745
title: "Завдання 364: між двома відрами"
challengeType: 28
dashedName: challenge-364
---

# --description--

Дано два відра фарби, кожне з кольором у форматі RGB і рівнем заповнення. Поверніть змішаний колір RGB у вигляді масиву з трьох цілих чисел.

- Кожне відро — це об’єкт (JavaScript) або словник (Python) з властивістю `color` (масив з трьох цілих чисел `[r, g, b]`) і властивістю `fullness` (від 0 до 100).
- Змішаний колір — це зважений середній кожного каналу двох кольорів на основі рівня заповнення, при цьому кожен канал округлюється до найближчого цілого числа.

# --hints--

`mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 })` має повернути `[125, 125, 125]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 50 }, { color: [0, 0, 0], fullness: 50 }), [125, 125, 125]);
```

`mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 })` має повернути `[200, 200, 200]`.

```js
assert.deepEqual(mixPaint({ color: [250, 250, 250], fullness: 80 }, { color: [0, 0, 0], fullness: 20 }), [200, 200, 200]);
```

`mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 })` має повернути `[100, 150, 200]`.

```js
assert.deepEqual(mixPaint({ color: [100, 150, 200], fullness: 30 }, { color: [100, 150, 200], fullness: 70 }), [100, 150, 200]);
```

`mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 })` має повернути `[114, 184, 170]`.

```js
assert.deepEqual(mixPaint({ color: [143, 143, 101], fullness: 45 }, { color: [100, 204, 204], fullness: 90 }), [114, 184, 170]);
```

`mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 })` має повернути `[68, 163, 123]`.

```js
assert.deepEqual(mixPaint({ color: [15, 134, 249], fullness: 29 }, { color: [97, 178, 55], fullness: 54 }), [68, 163, 123]);
```

# --seed--

## --seed-contents--

```js
function mixPaint(bucket1, bucket2) {

  return bucket1;
}
```

# --solutions--

```js
function mixPaint(bucket1, bucket2) {
  const total = bucket1.fullness + bucket2.fullness;
  return bucket1.color.map((channel, i) =>
    Math.round((channel * bucket1.fullness + bucket2.color[i] * bucket2.fullness) / total)
  );
}
```
