---
id: 6a26df3e2988bcdded204897
title: "Завдання 361: усний час"
challengeType: 28
dashedName: challenge-361
---

# --description--

За заданими кутами годинникової та хвилинної стрілок аналогового годинника в градусах (за годинниковою стрілкою від 12) поверніть час у вигляді усного англійського виразу.

Перетворіть кут хвилинної стрілки в хвилини (360° = 60 хвилин), а потім використайте такі правила:

| Хвилини | Усно |
|---------|--------|
| 0 | `"Y o'clock"` |
| 15 | `"quarter past Y"` |
| 1–29 (крім 15) | `"X minutes past Y"` |
| 30 | `"half past Y"` |
| 45 | `"quarter to Z"` |
| 31–59 (крім 45) | `"X minutes to Z"` (де `X` — 60 - хвилини) |

Де `Y` — поточна година, а `Z` — наступна година, обидві визначені за кутом годинникової стрілки (360° = 12 годин).

Примітка: кути стрілок можуть не збігатися точно з числом, тому варто якось їх округлити.

# --hints--

`getSpokenTime(90, 0)` має повернути `"3 o'clock"`.

```js
assert.equal(getSpokenTime(90, 0), "3 o'clock");
```

`getSpokenTime(160, 120)` має повернути `"20 minutes past 5"`.

```js
assert.equal(getSpokenTime(160, 120), "20 minutes past 5");
```

`getSpokenTime(255, 180)` має повернути `"half past 8"`.

```js
assert.equal(getSpokenTime(255, 180), "half past 8");
```

`getSpokenTime(67.5, 92)` має повернути `"quarter past 2"`.

```js
assert.equal(getSpokenTime(67.5, 92), "quarter past 2");
```

`getSpokenTime(200, 240)` має повернути `"20 minutes to 7"`.

```js
assert.equal(getSpokenTime(200, 240), "20 minutes to 7");
```

`getSpokenTime(322.5, 273)` має повернути `"quarter to 11"`.

```js
assert.equal(getSpokenTime(322.5, 273), "quarter to 11");
```

`getSpokenTime(117.5, 335)` має повернути `"5 minutes to 4"`.

```js
assert.equal(getSpokenTime(117.5, 335), "5 minutes to 4");
```

# --seed--

## --seed-contents--

```js
function getSpokenTime(hourAngle, minuteAngle) {

  return hourAngle;
}
```

# --solutions--

```js
function getSpokenTime(hourAngle, minuteAngle) {
  const minutes = Math.floor(minuteAngle / 6);
  const hour = Math.floor(hourAngle / 30) || 12;
  const nextHour = (hour % 12) + 1;

  if (minutes === 0) return `${hour} o'clock`;
  if (minutes === 15) return `quarter past ${hour}`;
  if (minutes === 30) return `half past ${hour}`;
  if (minutes === 45) return `quarter to ${nextHour}`;
  if (minutes < 30) return `${minutes} minutes past ${hour}`;
  return `${60 - minutes} minutes to ${nextHour}`;
}
```
