---
id: 615f38279e5c3d53692ea441
title: Schritt 7
challengeType: 0
dashedName: step-7
---

# --description--

Fassen Sie Ihre `h1`- und `p`-Elemente in einem `div`-Element zusammen. Geben Sie diesem `div` ein `class`-Attribut mit dem Wert `label`.

# --hints--

Sie sollten ein neues `div`-Element erstellen.

```js
assert(document.querySelector('div'));
```

Ihr neues `div`-Element sollte das `class`-Attribut mit dem Wert `label` haben.

```js
assert(document.querySelector('div')?.classList?.contains('label'));
```

Ihre `h1`- und `p`-Elemente sollten sich innerhalb Ihres neuen `.label`-Elements befinden.

```js
const children = [...document.querySelectorAll('h1'), ...document.querySelectorAll('p')];
assert(children?.every(child => child?.parentElement?.classList?.contains('label')));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  --fcc-editable-region--
  
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
  
  --fcc-editable-region--
</body>
</html>
```

```css
html {
  font-size: 16px;
}

body {
  font-family: 'Open Sans', sans-serif;
}
```
