---
id: 615f357957e370510f21ea16
title: Step 5
challengeType: 0
dashedName: step-5
---

# --description--

Erstellen Sie einen `body`-Selektor und geben Sie ihm eine `font-family`-Eigenschaft, die auf `Open Sans` gesetzt ist, mit einem Fallback von `sans-serif`.

Denken Sie daran, dass Schriftarten mit Leerzeichen im Namen in CSS in Anführungszeichen gesetzt werden müssen.

# --hints--

Sie sollten einen `body`-Selektor haben.

```js
assert(new __helpers.CSSHelp(document).getStyle('body'));
```

Ihr `body`-Selektor sollte eine `font-family`-Eigenschaft haben, die auf `"Open Sans", sans-serif` gesetzt ist.

```js
assert(new __helpers.CSSHelp(document).getStyle('body')?.fontFamily === '"Open Sans", sans-serif');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  <link href="https://fonts.googleapis.com/css?family=Open+Sans:400,700,800" rel="stylesheet">
  <link href="./styles.css" rel="stylesheet">
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
