---
id: 615f34948891834dd77655a6
title: Schritt 3
challengeType: 0
dashedName: step-3
---

# --description--

Fügen Sie ein zweites `p`-Element mit dem Text `Serving size 2/3 cup (55g)` hinzu.

# --hints--

Sie sollten ein zweites `p`-Element haben.

```js
assert(document.querySelectorAll('p')?.length === 2);
```

Ihr zweites `p`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert(document.querySelectorAll('p')?.[1]?.parentElement?.localName === 'body');
```

Ihr zweites `p`-Element sollte nach Ihrem vorhandenen `p`-Element kommen.

```js
assert(document.querySelectorAll('p')?.[1]?.previousElementSibling?.localName === 'p');
```

Ihr zweites `p`-Element sollte den Text `Serving size 2/3 cup (55g)` haben.

```js
assert(document.querySelectorAll('p')?.[1]?.innerText === 'Serving size 2/3 cup (55g)');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```

```css

```
