---
id: 615f34ecc1091b4fd5a8a484
title: Schritt 4
challengeType: 0
dashedName: step-4
---

# --description--

Fügen Sie innerhalb Ihres `head`-Elements ein `link`-Element mit dem `rel`-Attribut hinzu, das auf `stylesheet` gesetzt ist, und das `href`-Attribut, das auf `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` gesetzt ist.

Dadurch wird die `Open Sans`-Schriftfamilie importiert, mit den Schriftgewichtswerten `400`, `700` und `800`.

Fügen Sie außerdem ein `link`-Element hinzu, um Ihre `styles.css`-Datei zu verlinken.

# --hints--

Ihr Code sollte zwei `link`-Elemente enthalten.

```js
assert.strictEqual(document.querySelectorAll('link').length, 2);
```

Ihre `link`-Elemente sollten leere Elemente sein, sie dürfen kein End-Tag `</link>` haben.

```js
assert.notMatch(code, /<\/link>/);
```

Ihre zwei `link`-Elemente sollten sich innerhalb des `head`-Elements befinden.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.strictEqual(headElement.querySelectorAll('link').length, 2);
```

Beide Ihrer `link`-Elemente sollten das `rel`-Attribut mit dem Wert `stylesheet` haben.

```js
const links = [...document.querySelectorAll('link')];
assert.isTrue(links.every(link => link.getAttribute('rel') === 'stylesheet'));
```

Eines Ihrer `link`-Elemente sollte ein `href`-Attribut mit dem Wert `styles.css` haben.

```js
const styleElement = document.querySelector('[data-href]');
assert.isNotNull(styleElement);
```

Eines Ihrer `link`-Elemente sollte ein `href`-Attribut mit dem Wert `https://fonts.googleapis.com/css?family=Open+Sans:400,700,800` haben.

```js
const links = [...document.querySelectorAll('link')];
assert.exists(links.find(link => link?.getAttribute('href') === 'https://fonts.googleapis.com/css?family=Open+Sans:400,700,800'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
  --fcc-editable-region--
  
  --fcc-editable-region--
</head>

<body>
  <h1>Nutrition Facts</h1>
  <p>8 servings per container</p>
  <p>Serving size 2/3 cup (55g)</p>
</body>
</html>
```

```css

```
