---
id: 615f2d4150fe0d4cbd0f2628
title: Schritt 2
challengeType: 0
dashedName: step-2
---

# --description--

Fügen Sie unter Ihrem `h1`-Element ein `p`-Element mit dem Text `8 servings per container` hinzu.

# --hints--

Sie sollten ein neues `p`-Element hinzufügen.

```js
assert.exists(document.querySelector('p'));
```

Ihr `p`-Element sollte sich innerhalb Ihres `body`-Elements befinden.

```js
assert.equal(document.querySelector('p')?.parentElement?.localName, 'body');
```

Ihr `p`-Element sollte nach Ihrem `h1`-Element kommen.

```js
assert.equal(document.querySelector('p')?.previousElementSibling?.localName, 'h1');
```

Ihr `p`-Element sollte den Text `8 servings per container` enthalten.

```js
assert.equal(document.querySelector('p')?.innerText, '8 servings per container');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="UTF-8">
  <title>Nutrition Label</title>
</head>

<body>
  <h1>Nutrition Facts</h1>
  --fcc-editable-region--
  
  --fcc-editable-region--
</body>
</html>
```

```css

```
