---
id: a26cbbe9ad8655a977e1ceb5
title: Eine App zum Finden des längsten Wortes erstellen
challengeType: 26
dashedName: build-a-longest-word-finder-app
---

# --description--

In diesem Labor erstellen Sie eine Funktion, die die Länge des längsten Wortes im gegebenen Satz zurückgibt.

Zum Beispiel ist im Satz `"The quick brown fox jumped over the lazy dog"` das längste Wort `"jumped"`, das eine Länge von 6 hat.

**Ziel:** Erfüllen Sie die untenstehenden User Stories und bestehen Sie alle Tests, um das Labor abzuschließen.

**User Stories:**

1. Sie sollen eine Funktion namens `findLongestWordLength` erstellen, die einen String als Argument annimmt.
2. Die Funktion soll die Länge des längsten Wortes im String zurückgeben.

# --hints--

Sie sollen eine Funktion namens `findLongestWordLength` erstellen.

```js
assert.isFunction(findLongestWordLength);
```

`findLongestWordLength` soll einen einzelnen Parameter haben.

```js
assert.lengthOf(findLongestWordLength, 1);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` soll eine Zahl zurückgeben.

```js
assert.isNumber(
  findLongestWordLength('The quick brown fox jumped over the lazy dog')
);
```

`findLongestWordLength("The quick brown fox jumped over the lazy dog")` soll `6` zurückgeben.

```js
assert.strictEqual(
  findLongestWordLength('The quick brown fox jumped over the lazy dog'),
  6
);
```

`findLongestWordLength("May the force be with you")` soll `5` zurückgeben.

```js
assert.strictEqual(findLongestWordLength('May the force be with you'), 5);
```

`findLongestWordLength("Google do a barrel roll")` soll `6` zurückgeben.

```js
assert.strictEqual(findLongestWordLength('Google do a barrel roll'), 6);
```

`findLongestWordLength("Googling do a barrel roll")` soll `8` zurückgeben.

```js
assert.strictEqual(findLongestWordLength('Googling do a barrel roll'), 8);
```

`findLongestWordLength("What is the average airspeed velocity of an unladen swallow")` soll `8` zurückgeben.

```js
assert.strictEqual(
  findLongestWordLength(
    'What is the average airspeed velocity of an unladen swallow'
  ),
  8
);
```

`findLongestWordLength("What if we try a super-long word such as otorhinolaryngology")` soll `19` zurückgeben.

```js
assert.strictEqual(
  findLongestWordLength(
    'What if we try a super-long word such as otorhinolaryngology'
  ),
  19
);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function findLongestWordLength(str) {
  return str.split(' ').sort((a, b) => b.length - a.length)[0].length;
}

findLongestWordLength('The quick brown fox jumped over the lazy dog');
```
