---
id: 56533eb9ac21ba0edf2244cf
title: Costruire una collezione di dischi
challengeType: 26
dashedName: build-a-record-collection
---

# --description--

Stai creando una funzione che aiuta nella manutenzione di una collezione di album musicali. La collezione è organizzata come un oggetto che contiene più album, anch'essi oggetti. Ogni album è rappresentato nella collezione con un `id` unico come nome della proprietà. All'interno di ogni oggetto album, ci sono varie proprietà che descrivono informazioni sull'album. Non tutti gli album hanno informazioni complete.

La funzione `updateRecords` prende 4 argomenti rappresentati dai seguenti parametri della funzione:

-   `records` - un oggetto contenente diversi album individuali
-   `id` - un numero che rappresenta un album specifico nell'oggetto `records`
-   `prop` - una stringa che rappresenta il nome della proprietà dell'album da aggiornare
-   `value` - una stringa contenente le informazioni usate per aggiornare la proprietà dell'album

**Obiettivo:** Soddisfa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. La tua funzione deve sempre restituire l'intero oggetto `records`.
2. Se `value` è una stringa vuota, elimina la proprietà `prop` dall'album.
3. Se `prop` non è `tracks` e `value` non è una stringa vuota, assegna `value` alla proprietà `prop` di quell'album.
4. Se `prop` è `tracks` e `value` non è una stringa vuota, ma l'album non ha la proprietà `tracks`, crea un array vuoto e aggiungi `value` ad esso.
5. Se prop è `tracks` e `value` non è una stringa vuota, aggiungi `value` alla fine dell'array `tracks` esistente dell'album.

**Nota:** Per i test viene usata una copia dell'oggetto `recordCollection`. La tua funzione non dovrebbe fare riferimento diretto all'oggetto `recordCollection`, ma solo al parametro della funzione.

# --before-each--

```js
const _recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};
```

# --hints--

Dovresti avere una funzione `updateRecords`.

```js
assert.isFunction(updateRecords);
```

Dopo `updateRecords(recordCollection, 5439, "artist", "ABBA")`, `artist` dovrebbe essere la stringa `ABBA`

```js
assert.equal(
  updateRecords(_recordCollection, 5439, 'artist', 'ABBA')[5439]['artist'],
    'ABBA'
);
```

Dopo `updateRecords(recordCollection, 5439, "tracks", "Take a Chance on Me")`, `tracks` dovrebbe avere la stringa `Take a Chance on Me` come ultimo e unico elemento.

```js
updateRecords(_recordCollection, 5439, 'tracks', 'Take a Chance on Me');
assert.lengthOf(_recordCollection[5439]['tracks'], 1);
assert.equal(_recordCollection[5439]['tracks'].pop(), 'Take a Chance on Me');
```

Dopo `updateRecords(recordCollection, 2548, "artist", "")`, `artist` non dovrebbe essere impostato

```js
updateRecords(_recordCollection, 2548, 'artist', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('artist'));
```

Dopo `updateRecords(recordCollection, 1245, "tracks", "Addicted to Love")`, `tracks` dovrebbe avere la stringa `Addicted to Love` come ultimo elemento.

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'tracks', 'Addicted to Love')[1245][
    'tracks'
  ].pop(), 'Addicted to Love'
);
```

Dopo `updateRecords(recordCollection, 2468, "tracks", "Free")`, `tracks` dovrebbe avere la stringa `1999` come primo elemento.

```js
assert.equal(
  updateRecords(_recordCollection, 2468, 'tracks', 'Free')[2468][
    'tracks'
  ][0], '1999'
);
```

Dopo `updateRecords(recordCollection, 2548, "tracks", "")`, `tracks` non dovrebbe essere impostato

```js
updateRecords(_recordCollection, 2548, 'tracks', '');
assert.isFalse(_recordCollection[2548].hasOwnProperty('tracks'));
```

Dopo `updateRecords(recordCollection, 1245, "albumTitle", "Riptide")`, `albumTitle` dovrebbe essere la stringa `Riptide`

```js
assert.equal(
  updateRecords(_recordCollection, 1245, 'albumTitle', 'Riptide')[1245][
    'albumTitle'
  ], 'Riptide'
);
```

# --seed--

## --seed-contents--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};



```

# --solutions--

```js
const recordCollection = {
  2548: {
    albumTitle: 'Slippery When Wet',
    artist: 'Bon Jovi',
    tracks: ['Let It Rock', 'You Give Love a Bad Name']
  },
  2468: {
    albumTitle: '1999',
    artist: 'Prince',
    tracks: ['1999', 'Little Red Corvette']
  },
  1245: {
    artist: 'Robert Palmer',
    tracks: []
  },
  5439: {
    albumTitle: 'ABBA Gold'
  }
};

function updateRecords(records, id, prop, value) {
  if (value === '') delete records[id][prop];
  else if (prop === 'tracks') {
    records[id][prop] = records[id][prop] || [];
    records[id][prop].push(value);
  } else {
    records[id][prop] = value;
  }

  return records;
}
```
