---
id: 686daa7ed79ceacd0b264e80
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

점점 가까워지고 있습니다. 이제 이 예제를 살펴보세요:

```html
<svg viewBox="0 0 50 50">
</svg>
```

`viewBox` 속성은 SVG 안에서 이미지의 어떤 부분이 보일지 제어합니다.

- 처음 두 숫자(`0 0`)는 `viewBox`의 시작 위치, 즉 좌상단 모서리(x와 y)를 설정합니다.
- 다음 두 숫자(`50 50`)는 `viewBox`의 너비와 높이를 정의합니다.

`viewBox` 속성을 `0 0 24 24`으로 설정하세요.

# --hints--

`viewBox` 속성이 있어야 합니다.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
assert.exists(viewBox);
```

`viewBox` x 위치를 0으로 설정해야 합니다.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const x = viewBox?.trim().split(' ')[0];
assert.strictEqual(x, '0');
```

`viewBox` y 위치를 0으로 설정해야 합니다.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const y = viewBox?.trim().split(' ')[1];
assert.strictEqual(y, '0');
```

`viewBox` 너비를 24로 설정해야 합니다.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const width = viewBox?.trim().split(' ')[2];
assert.strictEqual(width, '24');
```

`viewBox` 높이를 24로 설정해야 합니다.

```js
const svg = document.querySelector('svg');
const viewBox = svg.getAttribute('viewBox');
const height = viewBox?.trim().split(' ')[3];
assert.strictEqual(height, '24');
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region--
    <svg width="24" height="24">
    --fcc-editable-region--
      <path
        d="M12 21s-6-4.35-9.33-8.22C-.5 7.39 3.24 1 8.4 4.28 10.08 5.32 12 7.5 12 7.5s1.92-2.18 3.6-3.22C20.76 1 24.5 7.39 21.33 12.78 18 16.65 12 21 12 21z"
      ></path>
    </svg>
  </body>
</html>
```
