---
id: 686d749d6f32f708fdf2b7ad
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이전 수업에서 `svg` 요소와 프로젝트에서 아이콘으로 자주 사용되는 방법을 배웠습니다. 실제 코드베이스에서는 보통 아이콘 라이브러리를 사용하므로 `svg` 요소를 처음부터 만들 필요가 없습니다. 하지만 이 워크숍에서는 `svg` 요소 안에서 사용되는 핵심 속성을 배우기 위해 하트 아이콘을 직접 만들어 보겠습니다.

페이지에 `svg` 요소를 만드는 것부터 시작하세요.

# --hints--

페이지에 `svg` 요소가 있어야 합니다.

```js
const svg = document.querySelector('svg');
assert.exists(svg);
```

# --seed--

## --seed-contents--

```html
<!doctype html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <title>Heart Icon</title>
  </head>
  <body>
    --fcc-editable-region-- 
    
    --fcc-editable-region--
  </body>
</html>
```
