---
id: 6140d1a351e88f159ed54fca
title: Step 10
challengeType: 0
dashedName: step-10
---

# --description--

同じパターンを使って、3番目の `.line`、4番目の `.line`、5番目の `.line`、6番目の `.line` に対して別々のセレクターを作成してください。

3番目の `transform` の `.line` プロパティを `rotate(120deg)` に、4番目を `rotate(180deg)` に、5番目を `rotate(240deg)` に、6番目を `rotate(300deg)` に設定してください。

# --hints--

`.line:nth-of-type(3)` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)'));
```

`.line:nth-of-type(3)` セレクターには `transform` プロパティを `rotate(120deg)` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(3)')?.transform, 'rotate(120deg)');
```

`.line:nth-of-type(4)` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)'));
```

`.line:nth-of-type(4)` セレクターには `transform` プロパティを `rotate(180deg)` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(4)')?.transform, 'rotate(180deg)');
```

`.line:nth-of-type(5)` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)'));
```

`.line:nth-of-type(5)` セレクターには `transform` プロパティを `rotate(240deg)` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(5)')?.transform, 'rotate(240deg)');
```

`.line:nth-of-type(6)` セレクターを作成する必要があります。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)'));
```

`.line:nth-of-type(6)` セレクターには `transform` プロパティを `rotate(300deg)` に設定してください。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.line:nth-of-type(6)')?.transform, 'rotate(300deg)');
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}

--fcc-editable-region--

--fcc-editable-region--
```
