---
id: 6140d0069049f5139d78da40
title: ステップ 8
challengeType: 0
dashedName: step-8
---

# --description--

`transform-origin` プロパティは、CSS変形が適用される基準点を設定するために使います。たとえば、このプロジェクトの後半で行う `rotate` 変形を適用するとき、`transform-origin` は要素がどの点を中心に回転するかを決めます。

`.line` セレクターに `transform-origin` プロパティを `0% 0%` に設定してください。これにより、基準点が左から `0%`、上から `0%` だけオフセットされ、要素の左上隅に設定されます。

# --hints--

`.line` セレクターには `transform-origin` プロパティを `0% 0%` に設定してください。

```js
const transformOrigin = new __helpers.CSSHelp(document).getStyle('.line')?.transformOrigin;
assert.oneOf(transformOrigin, ['0% 0%', '0% 0% 0px']);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

--fcc-editable-region--
.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
}
--fcc-editable-region--
```
