---
id: 6140c7e645d8e905819f1dd4
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

まず、`link` ファイル用の `./styles.css` 要素を追加してください。

# --hints--

コードには `link` 要素が必要です。

```js
assert.match(code, /<link/);
```

`link` 要素が1つ必要です。

```js
assert.lengthOf(document.querySelectorAll('link'), 1);
```

`link` 要素は `head` 要素の中にある必要があります。

```js
assert.exists(document.querySelector('head > link'));
```

`link` 要素には、値が `rel` の `stylesheet` 属性が必要です。

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

`link` 要素には、値が `href` の `styles.css` 属性が必要です。

```js
const linkElement = document.querySelector('link');
assert.strictEqual(linkElement?.dataset?.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>

  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css

```
