---
id: 6140d4bc9db3c81c51a09ab7
title: Step 15
challengeType: 0
dashedName: step-15
---

# --description--

パターンを続けて、次の `.cabin` 要素を選択し、それぞれに特定のルールを適用してください：

- 2番目の `.cabin` は `right` プロパティを `17%` に、`top` プロパティを `93.5%` に設定してください。
- 3番目の `.cabin` は `right` プロパティを `67%` に、`top` プロパティを `93.5%` に設定してください。
- 4番目の `.cabin` は `left` プロパティを `-8.5%` に、`top` プロパティを `50%` に設定してください。
- 5番目の `.cabin` は `left` プロパティを `17%` に、`top` プロパティを `7%` に設定してください。
- 6番目の `.cabin` は `right` プロパティを `17%` に、`top` プロパティを `7%` に設定してください。

# --hints--

`.cabin:nth-of-type(2)` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)'));
```

`.cabin:nth-of-type(2)` セレクターには `right` プロパティが `17%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.right, '17%');
```

`.cabin:nth-of-type(2)` セレクターには `top` プロパティが `93.5%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(2)')?.top, '93.5%');
```

`.cabin:nth-of-type(3)` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)'));
```

`.cabin:nth-of-type(3)` セレクターには `right` プロパティが `67%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.right,'67%');
```

`.cabin:nth-of-type(3)` セレクターには `top` プロパティが `93.5%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(3)')?.top, '93.5%');
```

`.cabin:nth-of-type(4)` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)'));
```

`.cabin:nth-of-type(4)` セレクターには `left` プロパティが `-8.5%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.left, '-8.5%');
```

`.cabin:nth-of-type(4)` セレクターには `top` プロパティが `50%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(4)')?.top, '50%');
```

`.cabin:nth-of-type(5)` セレクターがあるはずです。

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)'));
```

`.cabin:nth-of-type(5)` セレクターには `left` プロパティが `17%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.left, '17%');
```

`.cabin:nth-of-type(5)` セレクターには `top` プロパティが `7%` に設定されているはずです。

```js
assert.strictEqual(new __helpers.CSSHelp(document).getStyle('.cabin:nth-of-type(5)')?.top, '7%');
```

`.cabin:nth-of-type(6)` セレクターがあるはずです。

```js
const def = (s) => new __helpers.CSSHelp(document).getStyle(s)
assert.exists(def('.cabin:nth-of-type(6)') || def('.cabin:last-of-type'));
```

`.cabin:nth-of-type(6)` セレクターには `right` プロパティが `17%` に設定されているはずです。

```js
const right = (s) => new __helpers.CSSHelp(document).getStyle(s)?.right
assert.equal(right('.cabin:nth-of-type(6)') || right('.cabin:last-of-type'),'17%');
```

`.cabin:nth-of-type(6)` セレクターには `top` プロパティが `7%` に設定されているはずです。

```js
const top = (s) => new __helpers.CSSHelp(document).getStyle(s)?.top
assert.equal(top('.cabin:nth-of-type(6)') || top('.cabin:last-of-type'),'7%');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Ferris Wheel</title>
    <link rel="stylesheet" href="./styles.css">
  </head>
  <body>
    <div class="wheel">
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>
      <span class="line"></span>

      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
      <div class="cabin"></div>
    </div>
  </body>
</html>
```

```css
.wheel {
  border: 2px solid black;
  border-radius: 50%;
  margin-left: 50px;
  position: absolute;
  height: 55vw;
  width: 55vw;
  max-width: 500px;
  max-height: 500px;
}

.line {
  background-color: black;
  width: 50%;
  height: 2px;
  position: absolute;
  top: 50%;
  left: 50%;
  transform-origin: 0% 0%;
}

.line:nth-of-type(2) {
  transform: rotate(60deg);
}
.line:nth-of-type(3) {
  transform: rotate(120deg);
}
.line:nth-of-type(4) {
  transform: rotate(180deg);
}
.line:nth-of-type(5) {
  transform: rotate(240deg);
}
.line:nth-of-type(6) {
  transform: rotate(300deg);
}

.cabin {
  background-color: red;
  width: 20%;
  height: 20%;
  position: absolute;
  border: 2px solid;
  transform-origin: 50% 0%;
}

.cabin:nth-of-type(1) {
  right: -8.5%;
  top: 50%;
}
--fcc-editable-region--

--fcc-editable-region--
```
