---
id: 587d7fac367417b2b2512bde
title: "Verwenden Sie eine vordefinierte Skala, um Elemente zu platzieren"
challengeType: 6
forumTopicId: 301494
dashedName: use-a-pre-defined-scale-to-place-elements
---

# --description--

Nachdem die Skalen eingerichtet sind, ist es Zeit, das Streudiagramm erneut abzubilden. Die Skalen sind wie Verarbeitungsfunktionen, die die rohen `x` und `y` Daten in Werte umwandeln, die auf das SVG passen und korrekt gerendert werden. Sie halten die Daten innerhalb des Zeichenbereichs auf dem Bildschirm.

Sie setzen die Koordinaten-Attributwerte für eine SVG-Form mit der Skalierungsfunktion. Dazu gehören `x` und `y` Attribute für `rect` oder `text` Elemente oder `cx` und `cy` für `circles`. Hier ist ein Beispiel:

```js
shape.attr('x', d => xScale(d[0]));
```

Skalen setzen die Koordinatenattribute der Form, um die Datenpunkte auf dem SVG zu platzieren. Sie müssen die Skalen nicht anwenden, wenn Sie den tatsächlichen Datenwert anzeigen, zum Beispiel in der `text()` Methode für einen Tooltip oder eine Bezeichnung.

# --instructions--

Verwenden Sie `xScale` und `yScale`, um sowohl die `circle` als auch die `text` Formen auf dem SVG zu positionieren. Für die `circles` wenden Sie die Skalen an, um die `cx` und `cy` Attribute zu setzen. Geben Sie ihnen außerdem einen Radius von `5` Einheiten.

Für die `text` Elemente wenden Sie die Skalen an, um die `x` und `y` Attribute zu setzen. Die Bezeichnungen sollten rechts neben den Punkten versetzt sein. Fügen Sie dazu `10` Einheiten zum `x` Datenwert hinzu, bevor Sie ihn an die `xScale` übergeben.

# --hints--

Ihr Code sollte 10 `circle` Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

Das erste `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `91` und einen `cy` Wert von ungefähr `368` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const firstCircle = document.querySelectorAll('circle')[0];
assert.strictEqual(Math.round(firstCircle?.getAttribute('cx')), 91);
assert.strictEqual(Math.round(firstCircle?.getAttribute('cy')), 368);
assert.strictEqual(Math.round(firstCircle?.getAttribute('r')), 5);
```

Das zweite `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `159` und einen `cy` Wert von ungefähr `181` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const secondCircle = document.querySelectorAll('circle')[1];
assert.strictEqual(Math.round(secondCircle?.getAttribute('cx')), 159);
assert.strictEqual(Math.round(secondCircle?.getAttribute('cy')), 181);
assert.strictEqual(Math.round(secondCircle?.getAttribute('r')), 5);
```

Das dritte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `340` und einen `cy` Wert von ungefähr `329` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const thirdCircle = document.querySelectorAll('circle')[2];
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cx')), 340);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('cy')), 329);
assert.strictEqual(Math.round(thirdCircle?.getAttribute('r')), 5);
```

Das vierte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `131` und einen `cy` Wert von ungefähr `60` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const fourthCircle = document.querySelectorAll('circle')[3];
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('cy')), 60);
assert.strictEqual(Math.round(fourthCircle?.getAttribute('r')), 5);
```

Das fünfte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `440` und einen `cy` Wert von ungefähr `237` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const fifthCircle = document.querySelectorAll('circle')[4];
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cx')), 440);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('cy')), 237);
assert.strictEqual(Math.round(fifthCircle?.getAttribute('r')), 5);
```

Das sechste `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `271` und einen `cy` Wert von ungefähr `306` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const sixthCircle = document.querySelectorAll('circle')[5];
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cx')), 271);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('cy')), 306);
assert.strictEqual(Math.round(sixthCircle?.getAttribute('r')), 5);
```

Das siebte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `361` und einen `cy` Wert von ungefähr `351` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const seventhCircle = document.querySelectorAll('circle')[6];
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cx')), 361);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('cy')), 351);
assert.strictEqual(Math.round(seventhCircle?.getAttribute('r')), 5);
```

Das achte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `261` und einen `cy` Wert von ungefähr `132` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const eighthCircle = document.querySelectorAll('circle')[7];
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cx')), 261);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('cy')), 132);
assert.strictEqual(Math.round(eighthCircle?.getAttribute('r')), 5);
```

Das neunte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `131` und einen `cy` Wert von ungefähr `144` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const ninthCircle = document.querySelectorAll('circle')[8];
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cx')), 131);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('cy')), 144);
assert.strictEqual(Math.round(ninthCircle?.getAttribute('r')), 5);
```

Das zehnte `circle` Element sollte nach Anwendung der Skalen einen `cx` Wert von ungefähr `79` und einen `cy` Wert von ungefähr `326` haben. Es sollte außerdem einen `r` Wert von `5` besitzen.

```js
const tenthCircle = document.querySelectorAll('circle')[9];
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cx')), 79);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('cy')), 326);
assert.strictEqual(Math.round(tenthCircle?.getAttribute('r')), 5);
```

Ihr Code sollte 10 `text` Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Die erste Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `100` und einen `y` Wert von ungefähr `368` haben.

```js
const firstLabel = document.querySelectorAll('text')[0];
assert.strictEqual(Math.round(firstLabel?.getAttribute('x')), 100);
assert.strictEqual(Math.round(firstLabel?.getAttribute('y')), 368);
```

Die zweite Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `168` und einen `y` Wert von ungefähr `181` haben.

```js
const secondLabel = document.querySelectorAll('text')[1];
assert.strictEqual(Math.round(secondLabel?.getAttribute('x')), 168);
assert.strictEqual(Math.round(secondLabel?.getAttribute('y')), 181);
```

Die dritte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `350` und einen `y` Wert von ungefähr `329` haben.

```js
const thirdLabel = document.querySelectorAll('text')[2];
assert.strictEqual(Math.round(thirdLabel?.getAttribute('x')), 350);
assert.strictEqual(Math.round(thirdLabel?.getAttribute('y')), 329);
```

Die vierte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `141` und einen `y` Wert von ungefähr `60` haben.

```js
const fourthLabel = document.querySelectorAll('text')[3];
assert.strictEqual(Math.round(fourthLabel?.getAttribute('x')), 141);
assert.strictEqual(Math.round(fourthLabel?.getAttribute('y')), 60);
```

Die fünfte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `449` und einen `y` Wert von ungefähr `237` haben.

```js
const fifthLabel = document.querySelectorAll('text')[4];
assert.strictEqual(Math.round(fifthLabel?.getAttribute('x')), 449);
assert.strictEqual(Math.round(fifthLabel?.getAttribute('y')), 237);
```

Die sechste Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `280` und einen `y` Wert von ungefähr `306` haben.

```js
const sixthLabel = document.querySelectorAll('text')[5];
assert.strictEqual(Math.round(sixthLabel?.getAttribute('x')), 280);
assert.strictEqual(Math.round(sixthLabel?.getAttribute('y')), 306);
```

Die siebte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `370` und einen `y` Wert von ungefähr `351` haben.

```js
const seventhLabel = document.querySelectorAll('text')[6];
assert.strictEqual(Math.round(seventhLabel?.getAttribute('x')), 370);
assert.strictEqual(Math.round(seventhLabel?.getAttribute('y')), 351);
```

Die achte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `270` und einen `y` Wert von ungefähr `132` haben.

```js
const eighthLabel = document.querySelectorAll('text')[7];
assert.strictEqual(Math.round(eighthLabel?.getAttribute('x')), 270);
assert.strictEqual(Math.round(eighthLabel?.getAttribute('y')), 132);
```

Die neunte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `140` und einen `y` Wert von ungefähr `144` haben.

```js
const ninthLabel = document.querySelectorAll('text')[8];
assert.strictEqual(Math.round(ninthLabel?.getAttribute('x')), 140);
assert.strictEqual(Math.round(ninthLabel?.getAttribute('y')), 144);
```

Die zehnte Bezeichnung sollte nach Anwendung der Skalen einen `x` Wert von ungefähr `88` und einen `y` Wert von ungefähr `326` haben.

```js
const tenthLabel = document.querySelectorAll('text')[9];
assert.strictEqual(Math.round(tenthLabel?.getAttribute('x')), 88);
assert.strictEqual(Math.round(tenthLabel?.getAttribute('y')), 326);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1]);
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ', ' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));
  </script>
</body>
```
