---
id: 587d7fa7367417b2b2512bc8
title: Klassen mit D3 hinzufügen
challengeType: 6
forumTopicId: 301473
dashedName: add-classes-with-d3
---

# --description--

Viele inline Styles auf HTML-Elementen zu verwenden, wird schwer zu verwalten, selbst bei kleineren Apps. Es ist einfacher, den Elementen eine Klasse hinzuzufügen und diese Klasse einmal mit CSS-Regeln zu stylen. D3 hat die `attr()`-Methode, um jedem Element ein beliebiges HTML-Attribut hinzuzufügen, einschließlich eines Klassennamens.

Die `attr()`-Methode funktioniert genauso wie `style()`. Sie nimmt durch Kommas getrennte Werte und kann eine Callback-Funktion verwenden. Hier ist ein Beispiel, um einer Auswahl die Klasse `container` hinzuzufügen:

```js
selection.attr('class', 'container');
```

Beachten Sie, dass der `class`-Parameter immer gleich bleibt, wenn Sie eine Klasse hinzufügen müssen, und nur der `container`-Parameter sich ändert.

# --instructions--

Fügen Sie die `attr()`-Methode in den Code im Editor ein und geben Sie den `bar`-Elementen eine Klasse `div`.

# --hints--

Ihre `div`-Elemente sollten die Klasse `bar` haben.

```js
assert.isTrue(document.querySelector('div')?.classList.contains('bar'));
```

Ihr Code sollte die `attr()`-Methode verwenden.

```js
assert.match(code, /\.attr/g);
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line
  </script>
</body>
```
