---
id: 587d7fac367417b2b2512bdc
title: "Verwenden Sie die Funktionen d3.max und d3.min, um Minimum- und Maximumwerte in einem Datensatz zu finden"
challengeType: 6
forumTopicId: 301496
dashedName: >-
  use-the-d3-max-and-d3-min-functions-to-find-minimum-and-maximum-values-in-a-dataset
---

# --description--

Die D3-Methoden `domain()` und `range()` legen diese Informationen für Ihre Skala basierend auf den Daten fest. Es gibt einige Methoden, die das erleichtern.

Oft möchten Sie beim Festlegen des Bereichs das Minimum und Maximum innerhalb des Datensatzes verwenden. Diese Werte manuell zu finden, insbesondere bei großen Datensätzen, kann zu Fehlern führen.

D3 bietet zwei Methoden – `min()` und `max()` – um diese Informationen zurückzugeben. Hier ist ein Beispiel:

```js
const exampleData = [34, 234, 73, 90, 6, 52];
d3.min(exampleData);
d3.max(exampleData);
```

Ein Datensatz kann verschachtelte Arrays enthalten, wie die `[x, y]` Koordinatenpaare, die im Beispiel des Streudiagramms vorkamen. In diesem Fall müssen Sie D3 mitteilen, wie das Maximum und Minimum berechnet werden sollen. Glücklicherweise nehmen sowohl die `min()` als auch die `max()` Methoden eine Callback-Funktion entgegen. In diesem Beispiel ist das Argument `d` der Callback-Funktion das aktuelle innere Array. Die Callback-Funktion muss das Element aus dem inneren Array (den `x` oder `y` Wert) zurückgeben, über das Sie das Maximum oder Minimum berechnen möchten. Hier ein Beispiel, wie man die Min- und Max-Werte mit einem Array von Arrays findet:

```js
const locationData = [
  [1, 7],
  [6, 3],
  [8, 3]
];
const minX = d3.min(locationData, d => d[0]);
```

`minX` hätte den Wert `1`.

# --instructions--

Das `positionData` Array enthält Unterarrays mit x-, y- und z-Koordinaten. Verwenden Sie eine D3-Methode, um den Maximalwert der z-Koordinate (den dritten Wert) aus den Arrays zu finden und speichern Sie ihn in der Variablen `output`.

# --hints--

Der Text in `h2` sollte `8` sein.

```js
assert.strictEqual(output, 8);
assert.strictEqual(document.querySelector('h2')?.textContent, '8');
```

Ihr Code sollte die Methode `max()` verwenden.

```js
assert.match(
  code,
  /\.max/g,
  'Your code should use the <code>max()</code> method.'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];
    // Add your code below this line

    const output = undefined; // Change this line

    // Add your code above this line

    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const positionData = [
      [1, 7, -4],
      [6, 3, 8],
      [2, 9, 3]
    ];

    const output = d3.max(positionData, d => d[2]);

    d3.select('body').append('h2').text(output);
  </script>
</body>
```
