---
id: 587d7faa367417b2b2512bd4
title: Einen Hover-Effekt zu einem D3-Element hinzufügen
challengeType: 6
forumTopicId: 301469
dashedName: add-a-hover-effect-to-a-d3-element
---

# --description--

Es ist möglich, Effekte hinzuzufügen, die eine Leiste hervorheben, wenn der Benutzer mit der Maus darüber fährt. Bisher wird die Gestaltung der Rechtecke mit den eingebauten D3- und SVG-Methoden angewendet, aber Sie können auch CSS verwenden.

Sie setzen die CSS-Klasse auf die SVG-Elemente mit der `attr()`-Methode. Dann enthält die `:hover`-Pseudoklasse für Ihre neue Klasse die Stilregeln für alle Hover-Effekte.

# --instructions--

Verwenden Sie die `attr()`-Methode, um allen `bar`-Elementen eine Klasse `rect` hinzuzufügen. Dadurch ändert sich die `fill`-Farbe der Leiste zu Braun, wenn Sie mit der Maus darüber fahren.

# --hints--

Ihre `rect`-Elemente sollten eine Klasse `bar` haben.

```js
assert.isTrue(document.querySelector('rect')?.classList.contains('bar'));
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy')
    // Add your code below this line
      .attr('class', 'bar');
    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d - 3);
  </script>
</body>
```
