---
id: 587d7fac367417b2b2512bdd
title: Dynamische Skalen verwenden
challengeType: 6
forumTopicId: 301495
dashedName: use-dynamic-scales
---

# --description--

Die D3 `min()` und `max()` Methoden sind nützlich, um die Skala festzulegen.

Bei einem komplexen Datensatz ist eine Priorität, die Skala so einzustellen, dass die Visualisierung in die Breite und Höhe des SVG-Containers passt. Sie möchten, dass alle Daten innerhalb des SVG dargestellt werden, damit sie auf der Webseite sichtbar sind.

Das folgende Beispiel legt die x-Achsen-Skala für Streudiagrammdaten fest. Die `domain()` Methode übergibt der Skala Informationen über die Rohdatenwerte für das Diagramm. Die `range()` Methode gibt ihr Informationen über den tatsächlichen Platz auf der Webseite für die Visualisierung.

Im Beispiel reicht die Domäne von 0 bis zum Maximum im Datensatz. Sie verwendet die `max()` Methode mit einer Callback-Funktion, die auf den x-Werten in den Arrays basiert. Der Wertebereich nutzt die Breite des SVG (`w`), enthält aber auch etwas Innenabstand. Dadurch entsteht ein Abstand zwischen den Punkten des Streudiagramms und dem Rand des SVG.

```js
const dataset = [
  [34, 78],
  [109, 280],
  [310, 120],
  [79, 411],
  [420, 220],
  [233, 145],
  [333, 96],
  [222, 333],
  [78, 320],
  [21, 123]
];
const w = 500;
const h = 500;

const padding = 30;
const xScale = d3
  .scaleLinear()
  .domain([0, d3.max(dataset, d => d[0])])
  .range([padding, w - padding]);
```

Der Innenabstand kann anfangs verwirrend sein. Stellen Sie sich die x-Achse als eine horizontale Linie von 0 bis 500 vor (der Breitenwert des SVG). Wenn Sie den Innenabstand in der `range()` Methode berücksichtigen, beginnt das Diagramm bei 30 entlang dieser Linie (statt bei 0) und endet bei 470 (statt bei 500).

# --instructions--

Verwenden Sie die `yScale` Variable, um eine lineare y-Achsen-Skala zu erstellen. Die Domäne sollte bei null beginnen und bis zum maximalen `y` Wert im Datensatz reichen. Der Wertebereich sollte die Höhe des SVG (`h`) verwenden und Innenabstand enthalten.

**Hinweis:** Denken Sie daran, das Diagramm richtig herum zu halten. Wenn Sie den Wertebereich für die y-Koordinaten festlegen, ist der höhere Wert (Höhe minus Innenabstand) das erste Argument, und der niedrigere Wert das zweite Argument.

# --hints--

Der Text in der `h2` sollte `30` sein.

```js
assert.strictEqual(output, 30);
assert.strictEqual(document.querySelector('h2')?.textContent, '30');
```

Die `domain()` von yScale sollte `[0, 411]` entsprechen.

```js
assert.deepEqual(yScale.domain(), [0, 411]);
```

Die `range()` von yScale sollte `[470, 30]` entsprechen.

```js
assert.deepEqual(yScale.range(), [470, 30]);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    // Padding between the SVG boundary and the plot
    const padding = 30;

    // Create an x and y scale

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    // Add your code below this line

    const yScale = undefined;


    // Add your code above this line

    const output = yScale(411); // Returns 30
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const padding = 30;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const output = yScale(411);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
