---
id: 587d7fa8367417b2b2512bc9
title: Höhe eines Elements dynamisch aktualisieren
challengeType: 6
forumTopicId: 301493
dashedName: update-the-height-of-an-element-dynamically
---

# --description--

Die vorherigen Aufgaben behandelten, wie man Daten aus einem Array anzeigt und wie man CSS-Klassen hinzufügt. Sie können diese Lektionen kombinieren, um ein einfaches Balkendiagramm zu erstellen. Es gibt zwei Schritte dafür:

1. Erstellen Sie für jeden Datenpunkt im Array ein `div`

2. Geben Sie jedem `div` eine dynamische Höhe, indem Sie eine Callback-Funktion in der `style()`-Methode verwenden, die die Höhe auf den Datenwert setzt

Erinnern Sie sich an das Format, um einen Stil mit einer Callback-Funktion zu setzen:

```js
selection.style('cssProperty', d => d);
```

# --instructions--

Fügen Sie die `style()`-Methode zum Code im Editor hinzu, um die `height`-Eigenschaft für jedes Element zu setzen. Verwenden Sie eine Callback-Funktion, die den Wert des Datenpunkts mit dem angehängten String `px` zurückgibt.

# --hints--

Das erste `div` sollte eine `height` von `12` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[0]?.style.height, '12px');
```

Das zweite `div` sollte eine `height` von `31` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[1]?.style.height, '31px');
```

Das dritte `div` sollte eine `height` von `22` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[2]?.style.height, '22px');
```

Das vierte `div` sollte eine `height` von `17` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[3]?.style.height, '17px');
```

Das fünfte `div` sollte eine `height` von `25` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[4]?.style.height, '25px');
```

Das sechste `div` sollte eine `height` von `18` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[5]?.style.height, '18px');
```

Das siebte `div` sollte eine `height` von `29` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[6]?.style.height, '29px');
```

Das achte `div` sollte eine `height` von `14` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[7]?.style.height, '14px');
```

Das neunte `div` sollte eine `height` von `9` Pixeln haben.

```js
assert.strictEqual(document.querySelectorAll('div')[8]?.style.height, '9px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => `${d}px`);
  </script>
</body>
```
