---
id: 587d7fa8367417b2b2512bcc
title: Formen mit SVG anzeigen
challengeType: 6
forumTopicId: 301485
dashedName: display-shapes-with-svg
---

# --description--

Die letzte Aufgabe erstellte ein `svg`-Element mit einer vorgegebenen Breite und Höhe, das sichtbar war, weil ihm ein `background-color` im `style`-Tag zugewiesen wurde. Der Code schuf Platz für die vorgegebene Breite und Höhe.

Der nächste Schritt ist, eine Form zu erstellen, die in den `svg`-Bereich eingefügt wird. Es gibt verschiedene unterstützte Formen in SVG, wie Rechtecke und Kreise. Sie werden verwendet, um Daten darzustellen. Zum Beispiel könnte eine Rechteck-SVG-Form (`<rect>`) einen Balken in einem Balkendiagramm erzeugen.

Wenn Sie eine Form in den `svg`-Bereich platzieren, können Sie mit `x`- und `y`-Koordinaten angeben, wo sie sich befindet. Der Ursprungspunkt (0, 0) liegt in der oberen linken Ecke. Positive Werte für `x` verschieben die Form nach rechts, und positive Werte für `y` verschieben die Form vom Ursprungspunkt nach unten.

Um eine Form in der Mitte des 500 (Breite) x 100 (Höhe) `svg` aus der letzten Aufgabe zu platzieren, wäre die `x`-Koordinate 250 und die `y`-Koordinate 50.

Ein SVG-`rect` hat vier Attribute. Es gibt die `x`- und `y`-Koordinaten, die angeben, wo es im `svg`-Bereich platziert wird. Außerdem hat es ein `height`- und `width`, um die Größe zu bestimmen.

# --instructions--

Fügen Sie eine `rect`-Form zum `svg` mit `append()` hinzu und geben Sie ihr ein `width`-Attribut mit dem Wert `25` und ein `height`-Attribut mit dem Wert `100`. Geben Sie außerdem dem `rect` `x`- und `y`-Attribute, die jeweils auf `0` gesetzt sind.

# --hints--

Ihr Dokument sollte 1 `rect`-Element enthalten.

```js
assert.lengthOf(document.querySelectorAll('rect'), 1);
```

Das `rect`-Element sollte ein `width`-Attribut mit dem Wert `25` haben.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('width'), '25');
```

Das `rect`-Element sollte ein `height`-Attribut mit dem Wert `100` haben.

```js
assert.strictEqual(
  document.querySelector('rect')?.getAttribute('height'),
  '100'
);
```

Das `rect`-Element sollte ein `x`-Attribut mit dem Wert `0` haben.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('x'), '0');
```

Das `rect`-Element sollte ein `y`-Attribut mit dem Wert `0` haben.

```js
assert.strictEqual(document.querySelector('rect')?.getAttribute('y'), '0');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h)
      .append('rect')
      .attr('width', 25)
      .attr('height', 100)
      .attr('x', 0)
      .attr('y', 0);
  </script>
</body>
```
