---
id: 587d7fa6367417b2b2512bc3
title: Eine Gruppe von Elementen mit D3 auswählen
challengeType: 6
forumTopicId: 301490
dashedName: select-a-group-of-elements-with-d3
---

# --description--

D3 hat auch die `selectAll()`-Methode, um eine Gruppe von Elementen auszuwählen. Sie gibt ein Array von HTML-Knoten für alle Elemente im Dokument zurück, die mit dem Eingabestring übereinstimmen. Hier ist ein Beispiel, um alle Anker-Tags in einem Dokument auszuwählen:

```js
const anchors = d3.selectAll('a');
```

Wie die `select()`-Methode unterstützt `selectAll()` die Verkettung von Methoden, und Sie können sie mit anderen Methoden verwenden.

# --instructions--

Wählen Sie alle `li`-Tags im Dokument aus und ändern Sie deren Text in den String `list item`, indem Sie die `.text()`-Methode verketten.

# --hints--

Es sollten 3 `li`-Elemente auf der Seite sein, und der Text in jedem sollte genau `list item` lauten. Groß- und Kleinschreibung sowie Leerzeichen müssen exakt übereinstimmen.

```js
const listItems = document.querySelectorAll('li');
assert.lengthOf(listItems, 3);
for (let i = 0; i < listItems.length; i++) {
  assert.match(listItems[i]?.textContent, /list item/g);
}
```

Ihr Code sollte auf das `d3`-Objekt zugreifen.

```js
assert.match(code, /d3/g);
```

Ihr Code sollte die `selectAll`-Methode verwenden.

```js
assert.match(code, /\.selectAll/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <ul>
    <li>Example</li>
    <li>Example</li>
    <li>Example</li>
  </ul>
  <script>
    d3.selectAll('li').text('list item');
  </script>
</body>
```
