---
id: 587d7fa9367417b2b2512bd0
title: SVG-Elemente invertieren
challengeType: 6
forumTopicId: 301488
dashedName: invert-svg-elements
---

# --description--

Sie haben vielleicht bemerkt, dass das Balkendiagramm so aussieht, als wäre es auf dem Kopf stehend oder invertiert. Das liegt daran, wie SVG die (x, y)-Koordinaten verwendet.

Im SVG ist der Ursprungspunkt für die Koordinaten in der oberen linken Ecke. Eine `x`-Koordinate von 0 platziert eine Form am linken Rand des SVG-Bereichs. Eine `y`-Koordinate von 0 platziert eine Form am oberen Rand des SVG-Bereichs. Höhere `x`-Werte schieben das Rechteck nach rechts. Höhere `y`-Werte schieben das Rechteck nach unten.

Um die Balken richtig herum anzuzeigen, müssen Sie die Berechnung der `y`-Koordinate ändern. Sie muss sowohl die Höhe des Balkens als auch die Gesamthöhe des SVG-Bereichs berücksichtigen.

Die Höhe des SVG-Bereichs beträgt 100. Wenn Sie einen Datenpunkt mit dem Wert 0 im Satz haben, soll der Balken am unteren Rand des SVG-Bereichs beginnen (nicht oben). Dafür benötigt die `y`-Koordinate den Wert 100. Wenn der Datenpunkt den Wert 1 hätte, würden Sie mit einer `y`-Koordinate von 100 beginnen, um den Balken unten zu setzen. Dann müssen Sie die Höhe des Balkens von 1 berücksichtigen, sodass die endgültige `y`-Koordinate 99 wäre.

Die `y`-Koordinate, die `y = heightOfSVG - heightOfBar` ist, würde die Balken richtig herum platzieren.

# --instructions--

Ändern Sie die Callback-Funktion für das `y`-Attribut, damit die Balken richtig herum angezeigt werden. Denken Sie daran, dass die `height` des Balkens das Dreifache des Datenwerts `d` ist.

**Hinweis:** Allgemein gilt die Beziehung `y = h - m * d`, wobei `m` die Konstante ist, die die Datenpunkte skaliert.

# --hints--

Der erste `rect` sollte einen `y`-Wert von `64` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('y'),
  (h - dataset[0] * 3).toString()
);
```

Der zweite `rect` sollte einen `y`-Wert von `7` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('y'),
  (h - dataset[1] * 3).toString()
);
```

Der dritte `rect` sollte einen `y`-Wert von `34` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('y'),
  (h - dataset[2] * 3).toString()
);
```

Der vierte `rect` sollte einen `y`-Wert von `49` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('y'),
  (h - dataset[3] * 3).toString()
);
```

Der fünfte `rect` sollte einen `y`-Wert von `25` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('y'),
  (h - dataset[4] * 3).toString()
);
```

Der sechste `rect` sollte einen `y`-Wert von `46` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('y'),
  (h - dataset[5] * 3).toString()
);
```

Der siebte `rect` sollte einen `y`-Wert von `13` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('y'),
  (h - dataset[6] * 3).toString()
);
```

Der achte `rect` sollte einen `y`-Wert von `58` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('y'),
  (h - dataset[7] * 3).toString()
);
```

Der neunte `rect` sollte einen `y`-Wert von `73` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('y'),
  (h - dataset[8] * 3).toString()
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => {
        // Add your code below this line


        
        // Add your code above this line
      })
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
  </script>
</body>
```
