---
id: 587d7fa8367417b2b2512bcb
title: Lernen Sie SVG in D3 kennen
challengeType: 6
forumTopicId: 301489
dashedName: learn-about-svg-in-d3
---

# --description--

<dfn>SVG</dfn> steht für <dfn>Scalable Vector Graphics</dfn>.

Hier bedeutet „scalable“, dass ein Objekt beim Rein- oder Rauszoomen nicht pixelig erscheint. Es skaliert mit dem Anzeigesystem, egal ob auf einem kleinen Mobilbildschirm oder einem großen TV-Monitor.

SVG wird verwendet, um gängige geometrische Formen zu erstellen. Da D3 Daten in eine visuelle Darstellung umwandelt, nutzt es SVG, um die Formen für die Datenvisualisierung zu erzeugen. SVG-Formen für eine Webseite müssen innerhalb eines HTML-`svg`-Tags platziert werden.

CSS kann skalierbar sein, wenn Stile relative Einheiten verwenden (wie `vh`, `vw` oder Prozentsätze), aber SVG ist flexibler, um Datenvisualisierungen zu erstellen.

# --instructions--

Fügen Sie einen `svg`-Knoten zum `body` mit `append()` hinzu. Geben Sie ihm ein `width`-Attribut, das auf die bereitgestellte `w`-Konstante gesetzt ist, und ein `height`-Attribut, das auf die bereitgestellte `h`-Konstante gesetzt ist, wobei Sie jeweils die `attr()`- oder `style()`-Methode verwenden. Sie sehen es in der Ausgabe, weil ein `background-color` in Pink im `style`-Tag darauf angewendet wird.

**Hinweis:** Wenn Sie `attr()` verwenden, haben Breiten- und Höhenattribute keine Einheiten. Dies ist die Grundlage für das Skalieren – das Element hat immer ein Breiten-Höhen-Verhältnis von 5:1, egal wie stark gezoomt wird.

# --hints--

Ihr Dokument sollte 1 `svg`-Element enthalten.

```js
assert.lengthOf(document.querySelectorAll('svg'), 1);
```

Das `svg`-Element sollte ein `width`-Attribut mit dem Wert `500` haben oder so gestylt sein, dass es eine Breite von `500px` hat.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const width = window.getComputedStyle(svg)['width'];
assert.isTrue(svg?.getAttribute('width') === '500' || width === '500px');
```

Das `svg`-Element sollte ein `height`-Attribut mit dem Wert `100` haben oder so gestylt sein, dass es eine Höhe von `100px` hat.

```js
const svg = document.querySelector('svg');
assert.exists(svg); 
const height = window.getComputedStyle(svg)['height'];
assert.isTrue(svg?.getAttribute('height') === '100' || height === '100px');
```

# --seed--

## --seed-contents--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3.select('body');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<style>
  svg {
    background-color: pink;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);
  </script>
</body>
```
