---
id: 587d7fab367417b2b2512bd9
title: Bezeichnungen zu Streudiagramm-Kreisen hinzufügen
challengeType: 6
forumTopicId: 301477
dashedName: add-labels-to-scatter-plot-circles
---

# --description--

Sie können Text hinzufügen, um Bezeichnungen für die Punkte in einem Streudiagramm zu erstellen.

Das Ziel ist es, die durch Kommas getrennten Werte für das erste (`x`) und zweite (`y`) Feld jedes Elements in `dataset` anzuzeigen.

Die `text`-Knoten benötigen `x`- und `y`-Attribute, um sie auf dem SVG zu positionieren. In dieser Aufgabe kann der `y`-Wert (der die Höhe bestimmt) denselben Wert verwenden, den der `circle` für sein `cy`-Attribut nutzt. Der `x`-Wert kann etwas größer sein als der `cx`-Wert des `circle`, damit die Bezeichnung sichtbar ist. Dadurch wird die Bezeichnung rechts neben dem geplotteten Punkt angezeigt.

# --instructions--

Beschriften Sie jeden Punkt im Streudiagramm mit den `text`-Elementen. Der Text der Bezeichnung sollte die beiden Werte sein, getrennt durch ein Komma und ein Leerzeichen. Zum Beispiel ist die Bezeichnung für den ersten Punkt `34, 78`. Setzen Sie das `x`-Attribut so, dass es um `5` Einheiten größer ist als der Wert, den Sie für das `cx`-Attribut am `circle` verwendet haben. Setzen Sie das `y`-Attribut genauso, wie es für den `cy`-Wert am `circle` verwendet wird.

# --hints--

Ihr Code sollte 10 `text`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('text'), 10);
```

Die erste Bezeichnung sollte den Text `34, 78`, einen `x`-Wert von `39` und einen `y`-Wert von `422` haben.

```js
const labelOne = document.querySelector('text');
assert.strictEqual(labelOne?.textContent, '34, 78');
assert.strictEqual(labelOne?.getAttribute('x'), '39');
assert.strictEqual(labelOne?.getAttribute('y'), '422');
```

Die zweite Bezeichnung sollte den Text `109, 280`, einen `x`-Wert von `114` und einen `y`-Wert von `220` haben.

```js
const labelTwo = document.querySelectorAll('text')[1];
assert.strictEqual(labelTwo?.textContent, '109, 280');
assert.strictEqual(labelTwo?.getAttribute('x'), '114');
assert.strictEqual(labelTwo?.getAttribute('y'), '220');
```

Die dritte Bezeichnung sollte den Text `310, 120`, einen `x`-Wert von `315` und einen `y`-Wert von `380` haben.

```js
const labelThree = document.querySelectorAll('text')[2];
assert.strictEqual(labelThree?.textContent, '310, 120');
assert.strictEqual(labelThree?.getAttribute('x'), '315');
assert.strictEqual(labelThree?.getAttribute('y'), '380');
```

Die vierte Bezeichnung sollte den Text `79, 411`, einen `x`-Wert von `84` und einen `y`-Wert von `89` haben.

```js
const labelFour = document.querySelectorAll('text')[3];
assert.strictEqual(labelFour?.textContent, '79, 411');
assert.strictEqual(labelFour?.getAttribute('x'), '84');
assert.strictEqual(labelFour?.getAttribute('y'), '89');
```

Die fünfte Bezeichnung sollte den Text `420, 220`, einen `x`-Wert von `425` und einen `y`-Wert von `280` haben.

```js
const labelFive = document.querySelectorAll('text')[4];
assert.strictEqual(labelFive?.textContent, '420, 220');
assert.strictEqual(labelFive?.getAttribute('x'), '425');
assert.strictEqual(labelFive?.getAttribute('y'), '280');
```

Die sechste Bezeichnung sollte den Text `233, 145`, einen `x`-Wert von `238` und einen `y`-Wert von `355` haben.

```js
const labelSix = document.querySelectorAll('text')[5];
assert.strictEqual(labelSix?.textContent, '233, 145');
assert.strictEqual(labelSix?.getAttribute('x'), '238');
assert.strictEqual(labelSix?.getAttribute('y'), '355');
```

Die siebte Bezeichnung sollte den Text `333, 96`, einen `x`-Wert von `338` und einen `y`-Wert von `404` haben.

```js
const labelSeven = document.querySelectorAll('text')[6];
assert.strictEqual(labelSeven?.textContent, '333, 96');
assert.strictEqual(labelSeven?.getAttribute('x'), '338');
assert.strictEqual(labelSeven?.getAttribute('y'), '404');
```

Die achte Bezeichnung sollte den Text `222, 333`, einen `x`-Wert von `227` und einen `y`-Wert von `167` haben.

```js
const labelEight = document.querySelectorAll('text')[7];
assert.strictEqual(labelEight?.textContent, '222, 333');
assert.strictEqual(labelEight?.getAttribute('x'), '227');
assert.strictEqual(labelEight?.getAttribute('y'), '167');
```

Die neunte Bezeichnung sollte den Text `78, 320`, einen `x`-Wert von `83` und einen `y`-Wert von `180` haben.

```js
const labelNine = document.querySelectorAll('text')[8];
assert.strictEqual(labelNine?.textContent, '78, 320');
assert.strictEqual(labelNine?.getAttribute('x'), '83');
assert.strictEqual(labelNine?.getAttribute('y'), '180');
```

Die zehnte Bezeichnung sollte den Text `21, 123`, einen `x`-Wert von `26` und einen `y`-Wert von `377` haben.

```js
const labelTen = document.querySelectorAll('text')[9];
assert.strictEqual(labelTen?.textContent, '21, 123');
assert.strictEqual(labelTen?.getAttribute('x'), '26');
assert.strictEqual(labelTen?.getAttribute('y'), '377');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg.selectAll('text').data(dataset).enter().append('text');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', (d, i) => d[0])
      .attr('cy', (d, i) => h - d[1])
      .attr('r', 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .attr('x', d => d[0] + 5)
      .attr('y', d => h - d[1])
      .text(d => d[0] + ', ' + d[1]);
  </script>
</body>
```
