---
id: 587d7fa7367417b2b2512bc7
title: Stile basierend auf Daten ändern
challengeType: 6
forumTopicId: 301479
dashedName: change-styles-based-on-data
---

# --description--

D3 dreht sich um die Visualisierung und Darstellung von Daten. Wahrscheinlich möchten Sie die Gestaltung von Elementen basierend auf den Daten ändern.
Zum Beispiel möchten Sie einen Datenpunkt blau färben, wenn sein Wert kleiner als 20 ist, und ansonsten rot. Sie können eine Callback-Funktion in der `style()`-Methode verwenden und die bedingte Logik einbauen. Die Callback-Funktion nutzt den `d`-Parameter, um den Datenpunkt darzustellen:

```js
selection.style('color', d => {});
```

Die `style()`-Methode ist nicht darauf beschränkt, nur das `color` zu setzen – sie kann auch mit anderen CSS-Eigenschaften verwendet werden.

# --instructions--

Fügen Sie die `style()`-Methode im Code-Editor hinzu, um die `color` der `h2`-Elemente bedingt zu setzen. Schreiben Sie die Callback-Funktion so, dass sie rot zurückgibt, wenn der Datenwert kleiner als 20 ist, und ansonsten grün.

**Hinweis:** Sie können if-else-Logik oder den ternären Operator verwenden.

# --hints--

Das erste `h2` sollte eine `color` von rot haben.

```js
const headingTwoFirst = document.querySelector('h2');
assert.exists(headingTwoFirst); 
const color = window.getComputedStyle(headingTwoFirst)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Das zweite `h2` sollte eine `color` von grün haben.

```js
const headingTwoSecond = document.querySelectorAll('h2')[1];
assert.exists(headingTwoSecond); 
const color = window.getComputedStyle(headingTwoSecond)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Das dritte `h2` sollte eine `color` von grün haben.

```js
const headingTwoThird = document.querySelectorAll('h2')[2];
assert.exists(headingTwoThird); 
const color = window.getComputedStyle(headingTwoThird)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Das vierte `h2` sollte eine `color` von rot haben.

```js
const headingTwoFourth = document.querySelectorAll('h2')[3];
assert.exists(headingTwoFourth); 
const color = window.getComputedStyle(headingTwoFourth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Das fünfte `h2` sollte eine `color` von grün haben.

```js
const headingTwoFifth = document.querySelectorAll('h2')[4];
assert.exists(headingTwoFifth); 
const color = window.getComputedStyle(headingTwoFifth)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Das sechste `h2` sollte eine `color` von rot haben.

```js
const headingTwoSixth = document.querySelectorAll('h2')[5];
assert.exists(headingTwoSixth); 
const color = window.getComputedStyle(headingTwoSixth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Das siebte `h2` sollte eine `color` von grün haben.

```js
const headingTwoSeventh = document.querySelectorAll('h2')[6];
assert.exists(headingTwoSeventh); 
const color = window.getComputedStyle(headingTwoSeventh)['color'];
assert.strictEqual(color, 'rgb(0, 128, 0)');
```

Das achte `h2` sollte eine `color` von rot haben.

```js
const headingTwoEighth = document.querySelectorAll('h2')[7];
assert.exists(headingTwoEighth); 
const color = window.getComputedStyle(headingTwoEighth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

Das neunte `h2` sollte eine `color` von rot haben.

```js
const headingTwoNinth = document.querySelectorAll('h2')[8];
assert.exists(headingTwoNinth); 
const color = window.getComputedStyle(headingTwoNinth)['color'];
assert.strictEqual(color, 'rgb(255, 0, 0)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD');
    // Add your code below this line

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => d + ' USD')
      .style('color', d => (d < 20 ? 'red' : 'green'));
  </script>
</body>
```
