---
id: 587d7fac367417b2b2512bdb
title: Setzen Sie eine Domain und eine Range auf einer Skala
challengeType: 6
forumTopicId: 301491
dashedName: set-a-domain-and-a-range-on-a-scale
---

# --description--

Standardmﾃ､ﾃ殃g verwenden Skalen die Identitﾃ､tsbeziehung. Das bedeutet, der Eingabewert wird auf den Ausgabewert abgebildet. Skalen kﾃｶnnen jedoch viel flexibler und interessanter sein.

Angenommen, ein Datensatz hat Werte im Bereich von 50 bis 480. Dies sind die Eingabedaten fﾃｼr eine Skala, auch bekannt als <dfn>Domain</dfn>.

Sie mﾃｶchten diese Punkte entlang der `x`-Achse im SVG zwischen 10 Einheiten und 500 Einheiten abbilden. Dies sind die Ausgabedaten, auch bekannt als <dfn>Range</dfn>.

Die `domain()`- und `range()`-Methoden setzen diese Werte fﾃｼr die Skala. Beide Methoden nehmen ein Array mit mindestens zwei Elementen als Argument. Hier ist ein Beispiel:

```js
scale.domain([50, 480]);
scale.range([10, 500]);
scale(50);
scale(480);
scale(325);
scale(750);
d3.scaleLinear();
```

In der Reihenfolge wﾃｼrden die folgenden Werte in der Konsole angezeigt: `10`, `500`, `323.37` und `807.67`.

Beachten Sie, dass die Skala die lineare Beziehung zwischen den Domain- und Range-Werten verwendet, um zu bestimmen, was die Ausgabe fﾃｼr eine gegebene Zahl sein soll. Der Minimalwert in der Domain (50) wird auf den Minimalwert (10) in der Range abgebildet.

# --instructions--

Erstellen Sie eine Skala und setzen Sie deren Domain auf `[250, 500]` und die Range auf `[10, 150]`.

**Hinweis:** Sie kﾃｶnnen die `domain()`- und `range()`-Methoden an die `scale`-Variable anhﾃ､ngen.

# --hints--

Ihr Code sollte die `domain()`-Methode verwenden.

```js
assert.match(code, /\.domain/g);
```

Die `domain()` der `scale` sollte auf `[250, 500]` gesetzt werden.

```js
assert.deepEqual(scale.domain(), [250, 500]);
```

Ihr Code sollte die `range()`-Methode verwenden.

```js
assert.match(code, /\.range/g);
```

Die `range()` der `scale` sollte auf `[10, 150]` gesetzt werden.

```js
assert.deepEqual(scale.range(), [10, 150]);
```

Der Text in der `h2` sollte `-102` sein.

```js
assert.strictEqual(document.querySelector('h2')?.textContent, '-102');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    // Add your code below this line
    const scale = d3.scaleLinear();



    // Add your code above this line
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const scale = d3.scaleLinear();
    scale.domain([250, 500]);
    scale.range([10, 150]);
    const output = scale(50);
    d3.select('body').append('h2').text(output);
  </script>
</body>
```
