---
id: 587d7faa367417b2b2512bd6
title: Einen Tooltip zu einem D3-Element hinzufügen
challengeType: 6
forumTopicId: 301470
dashedName: add-a-tooltip-to-a-d3-element
---

# --description--

Ein Tooltip zeigt mehr Informationen über ein Element auf einer Seite, wenn der Benutzer mit der Maus über dieses Element fährt. Es gibt verschiedene Möglichkeiten, einen Tooltip zu einer Visualisierung hinzuzufügen. Diese Aufgabe verwendet das SVG `title`-Element.

`title` wird mit der `text()`-Methode kombiniert, um dynamisch Daten zu den Balken hinzuzufügen.

# --instructions--

Fügen Sie unter jedem `title`-Knoten ein `rect`-Element an. Rufen Sie dann die `text()`-Methode mit einer Callback-Funktion auf, damit der Text den Datenwert anzeigt.

# --hints--

Ihr Code sollte 9 `title`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('title'), 9);
```

Das erste `title`-Element sollte den Tooltip-Text `12` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[0]?.textContent, '12');
```

Das zweite `title`-Element sollte den Tooltip-Text `31` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[1]?.textContent, '31');
```

Das dritte `title`-Element sollte den Tooltip-Text `22` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[2]?.textContent, '22');
```

Das vierte `title`-Element sollte den Tooltip-Text `17` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[3]?.textContent, '17');
```

Das fünfte `title`-Element sollte den Tooltip-Text `25` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[4]?.textContent, '25');
```

Das sechste `title`-Element sollte den Tooltip-Text `18` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[5]?.textContent, '18');
```

Das siebte `title`-Element sollte den Tooltip-Text `29` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[6]?.textContent, '29');
```

Das achte `title`-Element sollte den Tooltip-Text `14` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[7]?.textContent, '14');
```

Das neunte `title`-Element sollte den Tooltip-Text `9` haben.

```js
assert.strictEqual(document.querySelectorAll('title')[8]?.textContent, '9');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar');
    // Add your code below this line



    // Add your code above this line

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar:hover {
    fill: brown;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => d * 3)
      .attr('fill', 'navy')
      .attr('class', 'bar')
      .append('title')
      .text(d => d);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d)
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - (d * 3 + 3));
  </script>
</body>
```
