---
id: 587d7fad367417b2b2512bdf
title: Achsen zu einer Visualisierung hinzufügen
challengeType: 6
forumTopicId: 301472
dashedName: add-axes-to-a-visualization
---

# --description--

Eine weitere Möglichkeit, das Streudiagramm zu verbessern, besteht darin, eine x-Achse und eine y-Achse hinzuzufügen.

D3 hat zwei Methoden, `axisLeft()` und `axisBottom()`, um die y-Achse bzw. die x-Achse zu rendern. Hier ist ein Beispiel, um die x-Achse basierend auf dem `xScale` aus den vorherigen Aufgaben zu erstellen:

```js
const xAxis = d3.axisBottom(xScale);
```

Der nächste Schritt ist, die Achse auf dem SVG zu rendern. Dazu können Sie eine allgemeine SVG-Komponente, das `g`-Element, verwenden. `g` steht für Gruppe. Im Gegensatz zu `rect`, `circle` und `text` ist eine Achse beim Rendern nur eine gerade Linie. Da es sich um eine einfache Form handelt, funktioniert die Verwendung von `g`. Der letzte Schritt ist, ein `transform`-Attribut anzuwenden, um die Achse an der richtigen Stelle im SVG zu positionieren. Andernfalls würde die Linie entlang des Rahmens des SVG gerendert und wäre nicht sichtbar. SVG unterstützt verschiedene Arten von `transforms`, aber zum Positionieren einer Achse wird `translate` benötigt. Wenn es auf das `g`-Element angewendet wird, verschiebt es die gesamte Gruppe um die angegebenen Werte nach rechts und unten. Hier ist ein Beispiel:

```js
const xAxis = d3.axisBottom(xScale);

svg
  .append('g')
  .attr('transform', 'translate(0, ' + (h - padding) + ')')
  .call(xAxis);
```

Der obige Code platziert die x-Achse am unteren Rand des SVG. Dann wird sie als Argument an die `call()`-Methode übergeben. Die y-Achse funktioniert auf die gleiche Weise, außer dass das `translate`-Argument in der Form `(x, 0)` vorliegt. Da `translate` ein String in der oben genannten `attr()`-Methode ist, können Sie Verkettung verwenden, um variable Werte für seine Argumente einzuschließen.

# --instructions--

Das Streudiagramm hat jetzt eine x-Achse. Erstellen Sie eine y-Achse in einer Variablen namens `yAxis` mit der `axisLeft()`-Methode. Rendern Sie dann die Achse mit einem `g`-Element. Stellen Sie sicher, dass Sie ein `transform`-Attribut verwenden, um die Achse um die Anzahl der Innenabstandseinheiten nach rechts und `0` Einheiten nach unten zu verschieben. Denken Sie daran, die Achse zu `call()`.

# --hints--

Ihr Code sollte die `axisLeft()`-Methode mit `yScale` als Argument verwenden.

```js
assert.match(code, /\.axisLeft\(yScale\)/g);
```

Das y-Achsen-`g`-Element sollte ein `transform`-Attribut haben, um die Achse um `(60, 0)` zu verschieben.

```js
const yAxis = document.querySelectorAll('g')[10];
assert.match(yAxis?.getAttribute('transform'), /translate\(60\s*,\s*0\)/g);
```

Ihr Code sollte die `yAxis` aufrufen.

```js
assert.match(code, /\.call\(\s*yAxis\s*\)/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);
    // Add your code below this line
    const yAxis = undefined;
    // Add your code above this line

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;
    const padding = 60;

    const xScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[0])])
      .range([padding, w - padding]);

    const yScale = d3
      .scaleLinear()
      .domain([0, d3.max(dataset, d => d[1])])
      .range([h - padding, padding]);

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('circle')
      .data(dataset)
      .enter()
      .append('circle')
      .attr('cx', d => xScale(d[0]))
      .attr('cy', d => yScale(d[1]))
      .attr('r', d => 5);

    svg
      .selectAll('text')
      .data(dataset)
      .enter()
      .append('text')
      .text(d => d[0] + ',' + d[1])
      .attr('x', d => xScale(d[0] + 10))
      .attr('y', d => yScale(d[1]));

    const xAxis = d3.axisBottom(xScale);

    const yAxis = d3.axisLeft(yScale);

    svg
      .append('g')
      .attr('transform', 'translate(0,' + (h - padding) + ')')
      .call(xAxis);

    svg
      .append('g')
      .attr('transform', 'translate(' + padding + ',0)')
      .call(yAxis);
  </script>
</body>
```
