---
id: 587d7fa9367417b2b2512bd1
title: Farbe eines SVG-Elements ändern
challengeType: 6
forumTopicId: 301480
dashedName: change-the-color-of-an-svg-element
---

# --description--

Die Balken sind an der richtigen Position, aber sie sind alle schwarz und gleichfarbig. SVG bietet eine Möglichkeit, die Farbe der Balken zu ändern.

In SVG wird eine `rect` Form mit dem `fill` Attribut eingefärbt. Es unterstützt Hex-Codes, Farbnamen und RGB-Werte sowie komplexere Optionen wie Verläufe und Transparenz.

# --instructions--

Fügen Sie eine `attr()` Methode hinzu, um die `fill` aller Balken auf die Farbe Navy zu setzen.

# --hints--

Alle Balken sollten eine `fill` Farbe Navy haben.

```js
const rect = document.querySelector('rect');
assert.exists(rect); 
const fill = window.getComputedStyle(rect)['fill'];
assert.strictEqual(fill, 'rgb(0, 0, 128)');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d);
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', (d, i) => h - 3 * d)
      .attr('width', 25)
      .attr('height', (d, i) => 3 * d)
      .attr('fill', 'navy');
  </script>
</body>
```
