---
id: 587d7fa9367417b2b2512bce
title: Dynamisch die Koordinaten für jede Säule festlegen
challengeType: 6
forumTopicId: 301487
dashedName: dynamically-set-the-coordinates-for-each-bar
---

# --description--

Die letzte Aufgabe erstellte und fügte für jeden Punkt in `svg` ein Rechteck zum `dataset`-Element hinzu, um eine Säule darzustellen. Leider lagen sie alle übereinander.

Die Platzierung eines Rechtecks wird durch die `x`- und `y`-Attribute gesteuert. Sie sagen D3, wo es im `svg`-Bereich mit dem Zeichnen der Form beginnen soll. Die letzte Aufgabe setzte beide auf 0, sodass jede Säule in der oberen linken Ecke platziert wurde.

Für ein Balkendiagramm sollten alle Säulen auf derselben vertikalen Ebene liegen, was bedeutet, dass der `y`-Wert für alle Säulen gleich bleibt (bei 0). Der `x`-Wert muss sich jedoch ändern, wenn Sie neue Säulen hinzufügen. Denken Sie daran, dass größere `x`-Werte die Elemente weiter nach rechts verschieben. Wenn Sie die Array-Elemente in `dataset` durchlaufen, sollte der `x`-Wert zunehmen.

Die `attr()`-Methode in D3 akzeptiert eine Callback-Funktion, um dieses Attribut dynamisch festzulegen. Die Callback-Funktion nimmt zwei Argumente entgegen, eines für den Datenpunkt selbst (normalerweise `d`) und eines für den Index des Datenpunkts im Array. Das zweite Argument für den Index ist optional. Hier ist das Format:

```js
selection.attr('property', (d, i) => {});
```

Es ist wichtig zu beachten, dass Sie keine `for`-Schleife schreiben oder `forEach()` verwenden müssen, um die Elemente im Datensatz zu durchlaufen. Erinnern Sie sich, dass die `data()`-Methode den Datensatz parst, und jede Methode, die nach `data()` verkettet wird, wird einmal für jedes Element im Datensatz ausgeführt.

# --instructions--

Ändern Sie die Callback-Funktion des `x`-Attributs so, dass sie den Index mal 30 zurückgibt.

**Hinweis:** Jede Säule hat eine Breite von 25, daher fügt das Erhöhen jedes `x`-Werts um 30 etwas Abstand zwischen den Säulen hinzu. Jeder Wert größer als 25 würde in diesem Beispiel funktionieren.

# --hints--

Die erste `rect` sollte einen `x`-Wert von `0` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('x'),
  '0'
);
```

Die zweite `rect` sollte einen `x`-Wert von `30` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('x'),
  '30'
);
```

Die dritte `rect` sollte einen `x`-Wert von `60` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('x'),
  '60'
);
```

Die vierte `rect` sollte einen `x`-Wert von `90` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('x'),
  '90'
);
```

Die fünfte `rect` sollte einen `x`-Wert von `120` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('x'),
  '120'
);
```

Die sechste `rect` sollte einen `x`-Wert von `150` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('x'),
  '150'
);
```

Die siebte `rect` sollte einen `x`-Wert von `180` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('x'),
  '180'
);
```

Die achte `rect` sollte einen `x`-Wert von `210` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('x'),
  '210'
);
```

Die neunte `rect` sollte einen `x`-Wert von `240` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('x'),
  '240'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        // Add your code below this line


      
        // Add your code above this line
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => {
        return i * 30;
      })
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
