---
id: 587d7fa7367417b2b2512bc5
title: Mit dynamischen Daten in D3 arbeiten
challengeType: 6
forumTopicId: 301498
dashedName: work-with-dynamic-data-in-d3
---

# --description--

Die letzten beiden Aufgaben behandeln die Grundlagen der dynamischen Anzeige von Daten mit D3 unter Verwendung der `data()`- und `enter()`-Methoden. Diese Methoden nehmen einen Datensatz und erzeugen zusammen mit der `append()`-Methode für jeden Eintrag im Datensatz ein neues DOM-Element.

In der vorherigen Aufgabe haben Sie für jedes Element im `h2`-Array ein neues `dataset`-Element erstellt, aber alle enthielten denselben Text, `New Title`. Das liegt daran, dass Sie die an jedes der `h2`-Elemente gebundenen Daten noch nicht verwendet haben.

Die D3 `text()`-Methode kann als Argument einen String oder eine Callback-Funktion annehmen:

```js
selection.text(d => d);
```

Im obigen Beispiel bezieht sich der Parameter `d` auf einen einzelnen Eintrag im Datensatz, an den eine Selektion gebunden ist.

Im aktuellen Beispiel ist das erste `h2`-Element an 12 gebunden, das zweite `h2`-Element an 31, das dritte `h2`-Element an 22 und so weiter.

# --instructions--

Ändern Sie die `text()`-Methode so, dass jedes `h2`-Element den entsprechenden Wert aus dem `dataset`-Array mit einem Leerzeichen und dem String `USD` anzeigt. Zum Beispiel sollte die erste Überschrift `12 USD` sein.

# --hints--

Die erste `h2` sollte den Text `12 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[0]?.textContent, '12 USD');
```

Die zweite `h2` sollte den Text `31 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[1]?.textContent, '31 USD');
```

Die dritte `h2` sollte den Text `22 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[2]?.textContent, '22 USD');
```

Die vierte `h2` sollte den Text `17 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[3]?.textContent, '17 USD');
```

Die fünfte `h2` sollte den Text `25 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[4]?.textContent, '25 USD');
```

Die sechste `h2` sollte den Text `18 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[5]?.textContent, '18 USD');
```

Die siebte `h2` sollte den Text `29 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[6]?.textContent, '29 USD');
```

Die achte `h2` sollte den Text `14 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[7]?.textContent, '14 USD');
```

Die neunte `h2` sollte den Text `9 USD` haben.

```js
assert.strictEqual(document.querySelectorAll('h2')[8]?.textContent, '9 USD');
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      // Add your code below this line

      .text('New Title');

    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('h2')
      .data(dataset)
      .enter()
      .append('h2')
      .text(d => `${d} USD`);
  </script>
</body>
```
