---
id: 587d7fa8367417b2b2512bca
title: Die Darstellung eines Balkendiagramms ändern
challengeType: 6
forumTopicId: 301481
dashedName: change-the-presentation-of-a-bar-chart
---

# --description--

Die letzte Aufgabe erstellte ein Balkendiagramm, aber es gibt ein paar Formatierungsänderungen, die es verbessern könnten:

1. Fügen Sie zwischen jedem Balken Abstand hinzu, um sie optisch zu trennen. Das erreichen Sie, indem Sie der CSS-Klasse `bar` einen Außenabstand hinzufügen.

2. Erhöhen Sie die Höhe der Balken, um die Unterschiede der Werte besser darzustellen. Das erreichen Sie, indem Sie den Wert mit einer Zahl multiplizieren, um die Höhe zu skalieren.

# --instructions--

Fügen Sie zuerst einen `margin` von `2px` zur `bar`-Klasse im `style`-Tag hinzu. Ändern Sie anschließend die Callback-Funktion in der `style()`-Methode so, dass sie einen Wert zurückgibt, der `10`-mal so groß ist wie der ursprüngliche Datenwert (plus dem `px`).

**Hinweis:** Wenn Sie jeden Datenpunkt mit derselben Konstanten multiplizieren, ändert sich nur die Skalierung. Es ist wie ein Zoom, und die Bedeutung der zugrunde liegenden Daten bleibt unverändert.

# --hints--

Der erste `div` sollte eine `height` von `120` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const firstDiv = document.querySelector('div');
assert.exists(firstDiv); 
const height = window.getComputedStyle(firstDiv)['height'];
const margin = window.getComputedStyle(firstDiv)['margin-right'];
assert.strictEqual(height, '120px');
assert.strictEqual(margin, '2px');
```

Der zweite `div` sollte eine `height` von `310` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const secondDiv = document.querySelectorAll('div')[1];
assert.exists(secondDiv); 
const height = window.getComputedStyle(secondDiv)['height'];
const margin = window.getComputedStyle(secondDiv)['margin-right'];
assert.strictEqual(height, '310px');
assert.strictEqual(margin, '2px');
```

Der dritte `div` sollte eine `height` von `220` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const thirdDiv = document.querySelectorAll('div')[2];
assert.exists(thirdDiv); 
const height = window.getComputedStyle(thirdDiv)['height'];
const margin = window.getComputedStyle(thirdDiv)['margin-right'];
assert.strictEqual(height, '220px');
assert.strictEqual(margin, '2px');
```

Der vierte `div` sollte eine `height` von `170` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const fourthDiv = document.querySelectorAll('div')[3];
assert.exists(fourthDiv); 
const height = window.getComputedStyle(fourthDiv)['height'];
const margin = window.getComputedStyle(fourthDiv)['margin-right'];
assert.strictEqual(height, '170px');
assert.strictEqual(margin, '2px');
```

Der fünfte `div` sollte eine `height` von `250` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const fifthDiv = document.querySelectorAll('div')[4];
assert.exists(fifthDiv); 
const height = window.getComputedStyle(fifthDiv)['height'];
const margin = window.getComputedStyle(fifthDiv)['margin-right'];
assert.strictEqual(height, '250px');
assert.strictEqual(margin, '2px');
```

Der sechste `div` sollte eine `height` von `180` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const sixthDiv = document.querySelectorAll('div')[5];
assert.exists(sixthDiv); 
const height = window.getComputedStyle(sixthDiv)['height'];
const margin = window.getComputedStyle(sixthDiv)['margin-right'];
assert.strictEqual(height, '180px');
assert.strictEqual(margin, '2px');
```

Der siebte `div` sollte eine `height` von `290` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const seventhDiv = document.querySelectorAll('div')[6];
assert.exists(seventhDiv); 
const height = window.getComputedStyle(seventhDiv)['height'];
const margin = window.getComputedStyle(seventhDiv)['margin-right'];
assert.strictEqual(height, '290px');
assert.strictEqual(margin, '2px');
```

Der achte `div` sollte eine `height` von `140` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const eighthDiv = document.querySelectorAll('div')[7];
assert.exists(eighthDiv); 
const height = window.getComputedStyle(eighthDiv)['height'];
const margin = window.getComputedStyle(eighthDiv)['margin-right'];
assert.strictEqual(height, '140px');
assert.strictEqual(margin, '2px');
```

Der neunte `div` sollte eine `height` von `90` Pixeln und einen `margin` von `2` Pixeln haben.

```js
const ninthDiv = document.querySelectorAll('div')[8];
assert.exists(ninthDiv); 
const height = window.getComputedStyle(ninthDiv)['height'];
const margin = window.getComputedStyle(ninthDiv)['margin-right'];
assert.strictEqual(height, '90px');
assert.strictEqual(margin, '2px');
```

# --seed--

## --seed-contents--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    /* Add your code below this line */

    
    /* Add your code above this line */
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d + 'px'); // Change this line
  </script>
</body>
```

# --solutions--

```html
<style>
  .bar {
    width: 25px;
    height: 100px;
    margin: 2px;
    display: inline-block;
    background-color: blue;
  }
</style>
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    d3.select('body')
      .selectAll('div')
      .data(dataset)
      .enter()
      .append('div')
      .attr('class', 'bar')
      .style('height', d => d * 10 + 'px');
  </script>
</body>
```
