---
id: 587d7fa9367417b2b2512bcf
title: Höhe jeder Leiste dynamisch ändern
challengeType: 6
forumTopicId: 301486
dashedName: dynamically-change-the-height-of-each-bar
---

# --description--

Die Höhe jeder Leiste kann auf den Wert des Datenpunkts im Array gesetzt werden, ähnlich wie der `x`-Wert dynamisch gesetzt wurde.

```js
selection.attr('property', (d, i) => {});
```

Hier wäre `d` der Wert des Datenpunkts, und `i` der Index des Datenpunkts im Array.

# --instructions--

Ändern Sie die Callback-Funktion für das `height`-Attribut so, dass sie den Datenwert mal 3 zurückgibt.

**Hinweis:** Denken Sie daran, dass das Multiplizieren aller Datenpunkte mit derselben Konstante die Daten skaliert (wie ein Zoom). Das hilft, die Unterschiede zwischen den Leistenwerten in diesem Beispiel besser zu erkennen.

# --hints--

Die erste `rect` sollte ein `height` von `36` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[0]?.getAttribute('height'),
  '36'
);
```

Die zweite `rect` sollte ein `height` von `93` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[1]?.getAttribute('height'),
  '93'
);
```

Die dritte `rect` sollte ein `height` von `66` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[2]?.getAttribute('height'),
  '66'
);
```

Die vierte `rect` sollte ein `height` von `51` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[3]?.getAttribute('height'),
  '51'
);
```

Die fünfte `rect` sollte ein `height` von `75` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[4]?.getAttribute('height'),
  '75'
);
```

Die sechste `rect` sollte ein `height` von `54` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[5]?.getAttribute('height'),
  '54'
);
```

Die siebte `rect` sollte ein `height` von `87` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[6]?.getAttribute('height'),
  '87'
);
```

Die achte `rect` sollte ein `height` von `42` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[7]?.getAttribute('height'),
  '42'
);
```

Die neunte `rect` sollte ein `height` von `27` haben.

```js
assert.strictEqual(
  document.querySelectorAll('rect')[8]?.getAttribute('height'),
  '27'
);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        // Add your code below this line


        
        
        // Add your code above this line
      });
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', (d, i) => i * 30)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', (d, i) => {
        return d * 3;
      });
  </script>
</body>
```
