---
id: 587d7fab367417b2b2512bd7
title: Ein Streudiagramm mit SVG-Kreisen erstellen
challengeType: 6
forumTopicId: 301484
dashedName: create-a-scatterplot-with-svg-circles
---

# --description--

Ein Streudiagramm ist eine weitere Art der Visualisierung. Es verwendet normalerweise Kreise, um Datenpunkte abzubilden, die jeweils zwei Werte haben. Diese Werte beziehen sich auf die `x`- und `y`-Achsen und werden verwendet, um den Kreis in der Visualisierung zu positionieren.

SVG hat ein `circle`-Tag, um die Kreisform zu erstellen. Es funktioniert ähnlich wie die `rect`-Elemente, die Sie für das Balkendiagramm verwendet haben.

# --instructions--

Verwenden Sie die `data()`-, `enter()`- und `append()`-Methoden, um `dataset` an neue `circle`-Elemente zu binden, die an das SVG angehängt werden.

**Hinweis:** Die Kreise sind noch nicht sichtbar, da wir ihre Attribute noch nicht gesetzt haben. Das machen wir in der nächsten Aufgabe.

# --hints--

Ihr Code sollte 10 `circle`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('circle'), 10);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle');
    // Add your code below this line



    // Add your code above this line
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [
      [34, 78],
      [109, 280],
      [310, 120],
      [79, 411],
      [420, 220],
      [233, 145],
      [333, 96],
      [222, 333],
      [78, 320],
      [21, 123]
    ];

    const w = 500;
    const h = 500;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg.selectAll('circle').data(dataset).enter().append('circle');
  </script>
</body>
```
