---
id: 587d7fa8367417b2b2512bcd
title: Erstellen Sie für jeden Datenpunkt im Satz eine Leiste
challengeType: 6
forumTopicId: 301482
dashedName: create-a-bar-for-each-data-point-in-the-set
---

# --description--

Die letzte Aufgabe fügte nur ein Rechteck zum `svg`-Element hinzu, um eine Leiste darzustellen. Hier kombinieren Sie, was Sie bisher über `data()`, `enter()` und SVG-Formen gelernt haben, um für jeden Datenpunkt in `dataset` ein Rechteck zu erstellen und anzuhängen.

Eine vorherige Aufgabe zeigte das Format, wie man für jedes Element in `div` ein `dataset` erstellt und anhängt:

```js
d3.select('body')?.selectAll('div').data(dataset).enter().append('div');
```

Es gibt einige Unterschiede bei der Arbeit mit `rect`-Elementen anstelle von `div`-Elementen. Die `rect`-Elemente müssen an ein `svg`-Element angehängt werden, nicht direkt an das `body`. Außerdem müssen Sie D3 mitteilen, wo jedes `rect` innerhalb des `svg`-Bereichs platziert werden soll. Die Platzierung der Leisten wird in der nächsten Aufgabe behandelt.

# --instructions--

Verwenden Sie die `data()`-, `enter()`- und `append()`-Methoden, um für jedes Element in `rect` ein `dataset` zu erstellen und anzuhängen. Die Leisten sollten alle übereinander angezeigt werden; dies wird in der nächsten Aufgabe korrigiert.

# --hints--

Ihr Dokument sollte 9 `rect`-Elemente enthalten.

```js
assert.lengthOf(document.querySelectorAll('rect'), 9);
```

Ihr Code sollte die `data()`-Methode verwenden.

```js
assert.match(code, /\.data/g);
```

Ihr Code sollte die `enter()`-Methode verwenden.

```js
assert.match(code, /\.enter/g);
```

Ihr Code sollte die `append()`-Methode verwenden.

```js
assert.match(code, /\.append/g);
```

# --seed--

## --seed-contents--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      // Add your code below this line

      // Add your code above this line
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```

# --solutions--

```html
<body>
  <script>
    const dataset = [12, 31, 22, 17, 25, 18, 29, 14, 9];

    const w = 500;
    const h = 100;

    const svg = d3
      .select('body')
      .append('svg')
      .attr('width', w)
      .attr('height', h);

    svg
      .selectAll('rect')
      .data(dataset)
      .enter()
      .append('rect')
      .attr('x', 0)
      .attr('y', 0)
      .attr('width', 25)
      .attr('height', 100);
  </script>
</body>
```
