---
id: 69008059d43d6751385883c1
title: "Eseguire il debug degli errori degli operatori di incremento e decremento in un'app con bug"
challengeType: 26
dashedName: debug-increment-and-decrement-operator-errors-in-a-buggy-app
---

# --description--

Un tuo amico ti ha chiesto di aiutarlo a eseguire il debug del suo codice. Sta esercitandosi con gli operatori di incremento e decremento ma ottiene risultati inattesi nella console.

La tua attività è aiutarlo a individuare gli errori e aggiornare il codice in modo che gli output della console corrispondano ai risultati attesi.

Puoi fare riferimento alle lezioni precedenti sugli operatori di incremento e decremento e sulla notazione prefisso e postfisso.

**Storie utente:**

1. Dovresti avere una variabile chiamata `updatedCounter` che produce correttamente il valore `11` usando l'operatore di incremento. Aggiorna l'espressione per usare la notazione prefisso o postfisso corretta invece di inserire il valore manualmente.
2. Dovresti avere una variabile chiamata `finalScore` che produce correttamente il valore `8` usando l'operatore di incremento. Aggiorna l'espressione per usare la notazione prefisso o postfisso corretta invece di inserire il valore manualmente.
3. Dovresti avere una variabile chiamata `updatedCoins` che produce correttamente il valore `2` usando l'operatore di decremento. Aggiorna l'espressione per usare la notazione prefisso o postfisso corretta invece di inserire il valore manualmente.
4. Dovresti avere una variabile chiamata `newHealth` che produce correttamente il valore `7` usando l'operatore di decremento. Aggiorna l'espressione per usare la notazione prefisso o postfisso corretta invece di inserire il valore manualmente.

# --hints--

Dovresti avere una variabile chiamata `updatedCounter` che usa la notazione di incremento prefisso (`++counter`).

```js
assert.match(code, /updatedCounter\s*=\s*\+\+counter\s*;/);
assert.strictEqual(typeof updatedCounter, 'number');
```

`updatedCounter` dovrebbe essere uguale a `11`.

```js
assert.strictEqual(updatedCounter, 11);
```

Dovresti avere una variabile chiamata `finalScore` che usa la notazione di incremento postfisso (`score++`).

```js
assert.match(code, /finalScore\s*=\s*score\+\+\s*;/);
assert.strictEqual(typeof finalScore, 'number');
```

`finalScore` dovrebbe essere uguale a `8`.

```js
assert.strictEqual(finalScore, 8);
```

Dovresti avere una variabile chiamata `updatedCoins` che usa la notazione di decremento prefisso (`--coins`).

```js
assert.match(code, /updatedCoins\s*=\s*--coins\s*;/);
assert.strictEqual(typeof updatedCoins, 'number');
```

`updatedCoins` dovrebbe essere uguale a `2`.

```js
assert.strictEqual(updatedCoins, 2);
```

Dovresti avere una variabile chiamata `newHealth` che usa la notazione di decremento postfisso (`health--`).

```js
assert.match(code, /newHealth\s*=\s*health--\s*;/);
assert.strictEqual(typeof newHealth, 'number');
```

`newHealth` dovrebbe essere uguale a `7`.

```js
assert.strictEqual(newHealth, 7);
```

# --seed--

## --seed-contents--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = counter++;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = ++score;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = coins--;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = --health;
console.log(`Health is still ${newHealth}`);
```

# --solutions--

```js
let counter = 10;
console.log(`Counter is currently ${counter}`);

let updatedCounter = ++counter;
console.log(`Counter is now ${updatedCounter}`);

let score = 8;
console.log(`Score is currently ${score}`);

let finalScore = score++;
console.log(`Score is still ${finalScore}`);

let coins = 3;
console.log(`Coins is currently ${coins}`);

let updatedCoins = --coins;
console.log(`Coins is now ${updatedCoins}`);

let health = 7;
console.log(`Health is currently ${health}`);

let newHealth = health--;
console.log(`Health is still ${newHealth}`);
```
