---
id: 65f01969115f933073b6be03
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

벡터는 길이(또는 크기)와 방향을 가진 객체로, 단일 숫자로 표현할 수 없습니다. 물리학에서 벡터는 힘, 속도, 가속도 및 기타 양을 나타내는 데 자주 사용됩니다.

이 프로젝트에서는 벡터 공간을 구축할 것입니다. 벡터 공간은 그 집합 내 요소들 사이에 일련의 연산이 정의된 집합입니다. 곧 모든 세부 사항을 배우게 됩니다.

우선, `Vector`라는 이름의 빈 클래스를 선언하여 프로젝트를 시작하세요.

# --hints--

`Vector`라는 이름의 빈 클래스를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_class("Vector")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
