---
id: 65f03d9f92eac9183a4d3281
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

벡터는 유클리드 평면에서 두 좌표 `x`와 `y`로 정의할 수 있습니다. 축의 원점과 점 `(x, y)` 사이의 거리가 벡터의 길이 또는 노름이 됩니다. 그리고 벡터의 방향은 `(x, y)`을 향하게 됩니다.

<img alt="a 2-dimensional vector of coordinates (2, 3)" src="https://cdn.freecodecamp.org/curriculum/python/2dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

`Vector` 클래스 내에 `__init__` 메서드를 만들고 세 개의 매개변수 `self`, `x`, `y`를 지정하세요.

# --hints--

`__init__` 클래스 안에 `Vector` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").has_function("__init__")`)) })
```

`__init__` 메서드는 세 개의 매개변수 `self`, `x`, `y`를 받아야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("Vector").find_function("__init__").has_args("self, x, y")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    pass
--fcc-editable-region--
```
