---
id: 65f056a405239e1dc4cc2854
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 `y`를 `y` 객체의 `Vector` 속성에 할당하세요.

# --hints--

`y`를 `self.y`에 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").find_variable("self.y").is_equivalent("self.y = y")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
--fcc-editable-region--
```
