---
id: 65f058e06f34fd1f0ee6e55d
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

다음으로, `Vector` 클래스 안에 빈 `norm` 메서드를 만들고 `self` 매개변수를 추가하세요.

# --hints--

`Vector` 클래스에는 `norm` 메서드가 있어야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("norm")
    `))
})
```

`norm` 메서드에는 `self` 매개변수가 있어야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("norm").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
```
