---
id: 65f055b9190fc41ca35549b8
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

Python은 이름에 앞뒤로 밑줄 두 개가 붙는 다양한 메서드를 제공합니다. `__init__`와 `__str__` 같은 것들을 이미 알고 있을 수 있습니다. 이러한 메서드는 `__<name>__` 명명 규칙을 따르며, 특별 메서드, 매직 메서드 또는 *dunder* (double underscore의 줄임말) 메서드라고 부릅니다.

클래스에 특별 메서드를 정의하면 그 클래스의 동작에 영향을 줍니다. 특정 상황에서 내부적으로 호출됩니다(예: 인스턴스 생성 시 `__init__`, 객체가 출력되거나 `__str__`에 전달될 때 `str()`). 이 프로젝트에서는 가장 자주 사용되는 특별 메서드 몇 가지를 배우게 됩니다.

지금은 `x`를 `x` 객체의 `Vector` 속성에 할당하세요.

# --hints--

`x`를 `self.x`에 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__init__").has_stmt("self.x = x")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        pass
--fcc-editable-region--
```
