---
id: 65f0694fb2296f3caadf8347
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

벡터 $\mathbf{a}$의 길이, 즉 노름은 일반적으로 $\\| \mathbf{a} \\|$로 표시합니다. 이는 각 성분의 제곱을 더한 값의 제곱근으로 계산할 수 있습니다:
\\[ \\| \mathbf{a} \\| = \sqrt{a_1^2 + a_2^2 + \ldots + a_n^2}
 \\]

벡터 노름을 계산하고 `norm` 메서드에서 결과를 반환하세요.

# --hints--

`norm` 메서드는 제공된 공식을 사용해 현재 벡터 인스턴스의 노름을 나타내는 숫자를 계산하고 반환해야 합니다. 값을 근사하지 마세요.

```js
({ test: () => runPython(`
v1 = Vector(2, 2.5)
v2 = Vector(0, -1)
v3 = Vector(-5, -0.3)
assert v1.norm() == 3.2015621187164243
assert v2.norm() == 1
assert v3.norm() == 5.008991914547277
`) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        pass
--fcc-editable-region--
```
