---
id: 65f06a8e5a57673d700c79c3
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

`Vector` 클래스 외부에서 `Vector`의 인스턴스를 생성하고 정수 `2`와 `3`을 `x`와 `y` 값으로 전달한 뒤 `v1`라는 변수에 할당하세요.

# --hints--

변수 `v1`를 선언하고 `Vector(2, 3)`를 할당해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_variable("v1").is_equivalent("v1 = Vector(2, 3)")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--
```
