---
id: 65f07c9b1ffb814d856dcffc
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

메서드가 반환하는 값을 출력하여 `norm`가 예상대로 작동하는지 테스트하세요.

# --hints--

`norm()`에서 `v1`을 호출하고 결과를 출력하세요.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
--fcc-editable-region--
```
