---
id: 65f3fe07cc763212efe91285
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

벡터는 여러 개의 차원(컴포넌트)을 가질 수 있습니다 `n`. 다음은 3차원 벡터의 표현입니다:

<img alt="a 3-dimensional vector of coordinates (4, 6, 3)" src="https://cdn.freecodecamp.org/curriculum/python/3dvector.png" style="background-color: white; height: 350px; width: auto; padding: 15px; display: block; margin-right: auto; margin-left: auto; margin-bottom: 1.2rem;">

지금까지 2차원 벡터를 만들었습니다. 각 경우마다 필요한 코드를 다시 작성하지 않고도 다양한 차원의 벡터를 나타낼 수 있길 원합니다. 이를 위해 상속을 사용할 것입니다.

먼저 `Vector` 클래스를 `R2Vector`로 이름을 바꾸어 이 클래스가 2차원 벡터를 다룬다는 것을 명시하세요. `v1`의 인스턴스화도 수정하는 것을 잊지 마세요.

# --hints--

`Vector` 클래스를 `R2Vector`로 이름을 바꾸세요. `v1`의 인스턴스화도 수정하는 것을 잊지 마세요.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

v1 = Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--
```
