---
id: 65f3fa60a93b84110b3f1708
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

의미 있는 문자열 표현을 반환하고 싶습니다. 예를 들어 `v1`의 경우, 벡터 컴포넌트 값이 포함된 튜플을 나타내는 문자열을 반환하고 싶습니다: `(2, 3)`.

`__str__` 메서드에서 벡터 컴포넌트를 순서대로 포함하는 튜플로 벡터를 나타내는 문자열을 반환하세요. 그런 다음 `Vector` 클래스 밖으로 나와 결과를 확인하기 위해 `v1`를 출력하세요.

# --hints--

`__str__` 메서드는 벡터 컴포넌트를 순서대로 포함하는 튜플로 벡터를 나타내는 문자열을 반환해야 합니다.

```js
({ test: () => runPython(`
v1 = Vector(2, 3)
v2 = Vector(0, 0)
v3 = Vector(-2, -3.5)
assert str(v1) == '(2, 3)'
assert str(v2) == '(0, 0)'
assert str(v3) == '(-2, -3.5)'
`) })
```

`v1`를 출력해야 합니다.

```js
({
    test: () => assert(runPython(`_Node(_code).has_call("print(v1)")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        pass

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
