---
id: 65f3fa097f26b510db6c710b
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

좋아요, `norm`가 예상대로 작동하고 있습니다. 이제 `v1`을 출력하려고 하면 객체의 기본 문자열 표현(예: `<__main__.Vector object at 0x11eb778>`)이 출력됩니다.

`Vector` 클래스 안에 읽기 쉬운 문자열 표현을 구현하기 위해 빈 `__str__` 메서드를 선언하세요. 매개변수 `self`를 반드시 지정해야 합니다.

`v1`이 문자열을 반환할 때까지 `__str__`을 출력하지 않도록 주의하세요. 그렇지 않으면 `TypeError`가 발생하는데, `__str__`는 항상 문자열을 반환해야 하기 때문입니다.

# --hints--

`__str__` 클래스 내에 `Vector` 메서드를 정의해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").has_function("__str__")
    `))
})
```

`__str__` 메서드는 `self` 매개변수를 가져야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("Vector").find_function("__str__").has_args("self")
    `))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5

v1 = Vector(2, 3)
print(v1.norm())
--fcc-editable-region--
```
