---
id: 65f4520e363e2642f8112e33
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

보시다시피, 뭔가 제대로 작동하지 않습니다. `norm`와 `__str__` 메서드는 `R2Vector`에서 상속받았지만 3차원 벡터에 적응하지 못합니다. 이들의 구현은 더 유연해야 합니다.

Python의 모든 객체는 `__dict__`라는 특별한 속성을 가지고 있는데, 이 속성은 객체 속성을 저장하는 사전입니다.

기존 `print` 호출을 제거하세요. 그런 다음 `__dict__`와 `v1` 벡터의 `v2` 속성을 출력해서 어떻게 생겼는지 확인하세요.

# --hints--

코드에 `print(v2)`와 `print(v2.norm())`가 없어야 합니다.

```js
({
    test: () => {
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2)")`));
      assert.isFalse(runPython(`_Node(_code).has_call("print(v2.norm())")`));
    }
})
```

`__dict__`와 `v1`의 `v2` 속성을 출력해야 합니다.

```js
({ test: () => assert(runPython(`
(_Node(_code).has_call("print(v1.__dict__)") and _Node(_code).has_call("print(v2.__dict__)")) or _Node(_code).has_call("print(v1.__dict__, v2.__dict__)")
`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v2.norm())
print(v2)
--fcc-editable-region--
```
