---
id: 65f40051d6b09a139f253e8e
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

상속을 사용하면 기존 클래스로부터 새로운 클래스를 정의할 수 있습니다. 새 클래스는 자식 클래스라고 하며, 기존 클래스인 부모 클래스의 모든 메서드와 프로프를 상속받습니다.

```py
class Tree:
    def sprout(self):
        print('Making new leaves!')

class Oak(Tree):
    pass
    
Oak().sprout() # Output: Making new leaves!
```

위 예제에서 자식 클래스 `Oak`는 `Tree`로부터 상속받아 부모 클래스 `sprout`의 `Tree` 메서드를 상속받습니다.

`R3Vector`라는 새 클래스를 만들고 위 예제를 따라 `R2Vector` 클래스로부터 상속받게 하세요.

# --hints--

새 클래스 `R3Vector`를 정의하세요.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_class("R3Vector")
    `))
})
```

새 클래스 `R3Vector`는 `R2Vector`로부터 상속받아야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").inherits_from("R2Vector")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--

--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
