---
id: 65f40fdb3579aa1ced28b2eb
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

`R3Vector` 인스턴스를 만들고 `2`, `2`, `3`를 각각 `x`, `y`, `z`의 값으로 사용하세요. 새 인스턴스를 `v2`에 할당하세요.

# --hints--

변수 `v2`를 선언하고 `R3Vector(2, 2, 3)`를 할당하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(2, 2, 3)")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
--fcc-editable-region--

--fcc-editable-region--
```
