---
id: 65f4535bbdb28d436ff3ddc9
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

출력에서 볼 수 있듯이 `__dict__`에는 인스턴스 속성 값들이 들어 있습니다. `self.x`과 `self.y`의 제곱을 명시적으로 더하는 대신, `__dict__`에 저장된 값들을 순회하여 노름 값을 계산할 것입니다.

`norm` 메서드 본문 내에서 소괄호 안 내용을 각 `val` 값을 `self.__dict__.values()`에서 `2` 제곱하는 제너레이터 표현식으로 바꾸세요. 또한, 그 제너레이터 표현식을 `sum` 함수의 인수로 전달하여 제곱값들이 모두 더해진 후 제곱근을 계산하도록 하세요.

# --hints--

`sum(val**2 for val in self.__dict__.values())**0.5` 메서드에서 `norm()`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in self.__dict__.values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
--fcc-editable-region--        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)

```
