---
id: 65f412c208c3791fee305acf
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

Python에서는 함수나 메서드 시그니처에 `*`를 추가 인수로 넣어 키워드 전용 인수 사용을 강제할 수 있습니다. `__init__` 메서드 둘 다 `*` 다음에 `self`를 두 번째 매개변수로 추가해 수정하세요. 그 이후에 오는 모든 매개변수는 함수나 메서드 호출 시 키워드 인수로 사용해야 합니다.

이것은 `super().__init__(x, y)` 호출도 수정해야 한다는 뜻입니다. `x`에 `x` 값을, `y`에 `y` 값을 할당하는 방식으로 하세요.

마지막으로 `v1`와 `v2` 인스턴스 생성 시 키워드 인수를 사용하도록 수정하세요.

# --hints--

두 `*` 메서드에서 두 번째 매개변수로 `__init__`를 추가해 키워드 인수를 강제하세요.

```js
({ test: () => {
  assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__init__").has_args("self, *, x, y")`));
  assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, *, x, y, z")`));
} })
```

`super.__init__(x, y)` 호출을 키워드 인수로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R3Vector").find_function("__init__").find_body().is_equivalent("super().__init__(x=x, y=y)\\nself.z = z")`)) })
```

`v1` 할당을 `R2Vector(x=2, y=3)`로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v1").is_equivalent("v1 = R2Vector(x=2, y=3)")`)) })
```

`v2` 할당을 `R3Vector(x=2, y=2, z=3)`로 수정하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("v2").is_equivalent("v2 = R3Vector(x=2, y=2, z=3)")`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        super().__init__(x, y)
        self.z = z

v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
v2 = R3Vector(2, 2, 3)
--fcc-editable-region--
```
