---
id: 65f40401e6ef53173c04e27d
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

새로운 클래스 안에 `__init__` 메서드를 선언하세요. 매개변수로 `self`, `x`, `y`, 그리고 `z`를 주세요.

# --hints--

`__init__` 클래스 안에 `R3Vector` 메서드를 정의해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").has_function("__init__")
    `))
})
```

`__init__` 메서드는 네 개의 매개변수 `self`, `x`, `y`, 그리고 `z`를 가져야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_args("self, x, y, z")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
