---
id: 65f407ea37ad6e181b90462e
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

이전처럼 각 `i` 매개변수를 `self.i`에 할당할 수 있습니다. 반복할 줄이 적긴 하지만, 이 반복을 피하는 한 가지 방법은 `super()` 함수를 사용하는 것입니다. `super()`는 부모 클래스를 암묵적으로 참조할 수 있게 해줍니다: `super().__init__(x, y)`는 부모 클래스의 `__init__` 메서드를 호출합니다.

`super().__init__(x, y)` 메서드에 `__init__` 호출을 추가하세요.

# --hints--

`super().__init__(x, y)` 메서드에 `__init__` 호출이 있어야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R3Vector").find_function("__init__").has_stmt("super().__init__(x, y)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, x, y):
        self.x = x
        self.y = y
        
    def norm(self):
        return (self.x**2 + self.y**2)**0.5
        
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, x, y, z):
        pass
--fcc-editable-region--
v1 = R2Vector(2, 3)
print(v1.norm())
print(v1)
```
