---
id: 65f8720685ec351abef26740
title: 24단계
challengeType: 20
dashedName: step-24
---

# --description--

적절한 문자열 표현을 얻으려면 `tuple()` 생성자를 호출하고 이전 단계에서 작성한 제너레이터 표현식을 인수로 전달하세요.

# --hints--

`__str__` 메서드가 반환한 표현식을 `tuple()` 생성자에 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("tuple(getattr(self, i) for i in vars(self))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return (getattr(self, i) for i in vars(self))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
