---
id: 65f876d17832001e8e1abb05
title: 28단계
challengeType: 20
dashedName: step-28
---

# --description--

변수 `arg_list`를 선언하세요. `key`의 각 키-값 쌍에 대해 `val`과 `vars(self).items()`를 순회하는 리스트 내포를 값으로 할당하고, 각 반복에서 f-문자열 `f'{key}={val}'`를 계산하세요.

# --hints--

변수 `arg_list`를 선언해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_variable("arg_list")`)) })
```

리스트 내포를 `vars(self).items()`에 할당해야 하며 `arg_list`를 순회합니다.

```js
({ test: () => runPython(`
import ast
var = _Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list")
assert isinstance(var.tree.value, ast.ListComp), "Expected arg_list to be a list comprehension"
assert (actual := var.find_comp_iters()[0].is_equivalent("vars(self).items()")), f"Expected vars(self).items(), got {actual}"
`) })
```

`arg_list`에 할당된 리스트 내포는 `key`와 `val`를 사용해 `vars(self).items()`를 순회해야 합니다.

```js
({ test: () => assert(runPython(`
_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_targets()[0].is_equivalent("key, val")
`)) })
```

`arg_list`에 할당된 리스트 내포는 `f'{key}={val}'`의 각 키-값 쌍에 대해 문자열 `vars(self).items()`를 계산해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").find_variable("arg_list").find_comp_expr().is_equivalent("f'{key}={val}'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
