---
id: 65f8749b13774b1d2e4a7fba
title: 27단계
challengeType: 20
dashedName: step-27
---

# --description--

`__str__` 메서드는 객체를 사람이 읽을 수 있는 문자열로 나타내지만, `__repr__` 메서드는 객체를 생성하는 데 필요한 문자열을 반환해야 합니다.

`__repr__` 클래스 내에 `self` 매개변수를 가진 `R2Vector` 메서드를 추가하세요.

# --hints--

`R2Vector` 클래스에 `__repr__` 메서드가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__repr__")`)) })
```

`__repr__` 메서드에 `self` 매개변수가 있어야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_args("self")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
