---
id: 65f870003444fb1a2ad171f2
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

두 개의 `print` 호출을 수정하여 벡터의 노름을 출력하고 `norm()` 메서드가 제대로 작동하는지 확인하세요.

# --hints--

`v1.norm()`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(v1.norm())")`)) })
```

`v2.norm()`를 출력해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(v2.norm())")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5

    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.__dict__)
print(v2.__dict__)
--fcc-editable-region--
```
