---
id: 65f8709620b2ce1a62608f5a
title: 23단계
challengeType: 20
dashedName: step-23
---

# --description--

문자열 입력에서 시작해 동적으로 몇몇 속성에 접근해야 할 때, 내장 `getattr()` 함수가 필요합니다. 이 함수는 첫 번째 인수로 객체를 받고, 두 번째 인수로 속성 이름을 담은 문자열을 받습니다.

`__str__` 메서드를 고치기 시작하세요. `__str__()`가 반환하는 문자열을 객체 속성을 반복하고 각 속성에 대해 `getattr()`를 호출하는 제너레이터 표현식으로 바꾸세요 `i`.

# --hints--

객체 속성 `vars(self)`를 반복하는 제너레이터 표현식을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_iters()[0].is_equivalent("vars(self)")`)) })
```

반복 변수로 `i`를 사용하는 제너레이터 표현식을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_targets()[0].is_equivalent("i")`)) })
```

`getattr(self, i)`의 각 항목 `i`에 대해 `vars(self)`를 호출하는 제너레이터 표현식을 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").find_return().find_comp_expr().is_equivalent("getattr(self, i)")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
