---
id: 65f872a0fe6aa21b456ad4fe
title: 25단계
challengeType: 20
dashedName: step-25
---

# --description--

마지막으로 `tuple()` 호출을 `str()` 함수에 인수로 전달하세요. 

# --hints--

현재 `tuple()` 메서드에서 반환하는 `__str__` 호출을 `str()` 함수에 전달해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__str__").has_return("str(tuple(getattr(self, i) for i in vars(self)))")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5
--fcc-editable-region--
    def __str__(self):
        return tuple(getattr(self, i) for i in vars(self))
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
