---
id: 65f9b486989cb90ff3e77ac8
title: 39단계
challengeType: 20
dashedName: step-39
---

# --description--

`__add__` 메서드의 `other` 매개변수는 덧셈 연산에서 `+` 연산자의 오른쪽에 위치한 피연산자를 나타냅니다.

오른쪽 피연산자가 왼쪽 피연산자와 같은 클래스의 객체가 아닌지 확인하려고 합니다(즉, `self`). 피연산자의 타입을 확인하여 이를 검증하는 `if` 문을 만드세요.

# --hints--

`if` 문을 만들어 `type(self)`와 `type(other)`가 다른지 확인해야 합니다.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
