---
id: 65f93b67169a9c703264458a
title: 30단계
challengeType: 20
dashedName: step-30
---

# --description--

메서드는 상속받았을 때 `R2Vector`와 `R3Vector` 객체를 생성하는 문자열을 반환해야 하므로, 클래스 이름을 직접 지정하여 문자열을 만들 수 없습니다.

클래스 이름은 `__class__.__name__`로 접근할 수 있습니다. `return` 메서드에 `__repr__` 문장을 추가하고 객체를 생성하는 데 필요한 문자열을 반환하세요.

# --hints--

벡터 객체를 생성하는 데 필요한 f-문자열을 만들고 `__class__.__name__` 메서드에서 반환하려면 `__repr__`를 사용하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__repr__").has_return("f'{self.__class__.__name__}({args})'")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))
--fcc-editable-region--
    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
