---
id: 65f9cb678070ca1668898c70
title: 44단계
challengeType: 20
dashedName: step-44
---

# --description--

두 객체를 더할 때 `__add__`가 내부적으로 호출되는 것과 같은 방식으로, 뺄셈의 경우에는 `__sub__` 메서드가 암묵적으로 호출됩니다.

이제 빈 `__sub__` 메서드를 정의하고 두 개의 매개변수 `self`와 `other`를 받으세요. 새 메서드 안에서 `self`와 `other`가 같은 클래스에 속하지 않는지 확인하는 if 문을 만들고, 이전과 같이 `NotImplemented`를 반환하세요.

# --hints--

`__sub__` 클래스 안에 `R2Vector` 메서드를 정의해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__sub__")
    `))
})
```

`__sub__` 메서드는 두 개의 매개변수 `self`과 `other`를 가져야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").has_args("self, other")
    `))
})
```

`if` 문을 만들어 `self`와 `other`가 같은 클래스에 속하지 않는지 확인해야 합니다.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_conditions()[0]
node.is_equivalent("type(self) != type(other)") or node.is_equivalent("type(other) != type(self)")
    `))
})
```

`NotImplemented` 본문 안에서 `if`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
