---
id: 65f9c2c2006feb1428ad2d4c
title: 43단계
challengeType: 20
dashedName: step-43
---

# --description--

`**` 연산자를 사용해 딕셔너리를 키워드 인수로 풀 수 있습니다.

```py
def spam(a, b):
    return a + b

my_dict = {a: 11, b: 4}

spam(**my_dict) # 15
```

`__class__`을 사용하고 `**kwargs`를 인수로 전달하여 현재 클래스의 인스턴스를 반환하세요.

# --hints--

`self.__class__(**kwargs)` 메서드에서 `__add__`를 반환해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").has_return("self.__class__(**kwargs)")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
