---
id: 65f9b17a638f0d0dcce8c354
title: 38단계
challengeType: 20
dashedName: step-38
---

# --description--

벡터 공간을 만들려면 벡터가 여러 경우에 어떻게 동작해야 하는지 정의해야 합니다. 벡터는 더해져서 새로운 벡터를 만들 수 있습니다.

특수 메서드 `__add__`를 구현하여 `+` 연산자를 사용해 두 객체를 더할 때 기본적으로 일어나는 동작을 재정의할 수 있습니다.

현재는 `R2Vector` 또는 `R3Vector`의 두 인스턴스를 더하려고 하면 예외가 발생합니다. `__add__` 클래스 내에 빈 `R2Vector` 메서드를 만들고, 두 개의 매개변수 `self`와 `other`를 받도록 하세요.

# --hints--

`__add__` 클래스 내에 `R2Vector` 메서드를 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").has_function("__add__")`)) })
```

`__add__` 메서드는 두 개의 매개변수 `self`과 `other`를 가져야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").has_args("self, other")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
