---
id: 65f9dd6e5a08af19c196c2df
title: 45단계
challengeType: 20
dashedName: step-45
---

# --description--

한 벡터에서 다른 벡터를 뺀 결과 벡터는 각 컴포넌트의 차이를 계산하여 얻습니다. 예를 들어, `(2, 4)`에서 `(7, 3)`을 빼면 세 번째 벡터 `(5, -1)`가 생성되는데, 여기서 `5`는 `7`와 `2`의 차이이고 `-1`는 `3`와 `4`의 차이입니다.

이전에 `__add__` 메서드 안에서 했던 것처럼, `if` 문장 뒤에 `kwargs` 변수를 선언하세요. 두 피연산자의 `i` 속성 값 차이를 값으로 하고, 현재 반복 중인 키와 같은 키를 가지는 키-값 쌍을 생성하는 객체 속성을 순회하는 딕셔너리 내포를 할당하세요.

# --hints--

`kwargs` 변수를 선언하고 객체 속성을 순회하는 딕셔너리 내포를 할당해야 합니다.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_iters()[0]
node.is_equivalent("vars(self)") or node.is_equivalent("self.__dict__")`)) })
```

`kwargs`에 할당된 딕셔너리 내포는 `i` 변수를 사용해 객체 속성을 순회해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

`kwargs`에 할당된 딕셔너리 내포는 `i`를 키로 사용해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_key().is_equivalent("i")`)) })
```

`kwargs`에 할당된 딕셔너리 내포는 각 `i`에 대해 `self`의 `i` 속성과 `other`의 `i` 속성 차이를 `i` 키에 할당해야 합니다.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__sub__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) - getattr(other, i)") or node.is_equivalent("self.__getattribute__(i) - other.__getattribute__(i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
