---
id: 65f9bc10a9fb1612e066e9e8
title: 42단계
challengeType: 20
dashedName: step-42
---

# --description--

두 벡터를 더할 때, 한 벡터의 각 컴포넌트는 다른 벡터의 같은 컴포넌트에 더해집니다. 예를 들어, `(1, 2)`와 `(2, 4)`을 더하면 세 번째 벡터 `(3, 6)`가 생성되는데, 여기서 `3`는 `1`와 `2`의 합이고 `6`는 `2`와 `4`의 합입니다.

`kwargs` 딕셔너리는 더해진 두 벡터와 같은 클래스의 새 벡터를 생성하는 데 필요한 키-값 쌍을 포함할 것입니다.

빈 딕셔너리를 `vars(self)`을 순회하는 딕셔너리 내포로 바꾸고, 각 키(`i`)에 대해 현재 반복의 키와 같고 그 값은 두 피연산자의 `i` 속성 값의 합인 키-값 쌍을 만드세요.

# --hints--

`kwargs`에 할당된 빈 딕셔너리를 `vars(self)`를 순회하는 딕셔너리 내포로 바꾸세요.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_iters()[0].is_equivalent("vars(self)")
    `))
})
```

`kwargs`에 할당된 딕셔너리 내포는 `i`를 순회하기 위해 `vars(self)` 변수를 사용해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_targets()[0].is_equivalent("i")
    `))
})
```

`kwargs`에 할당된 딕셔너리 내포는 키로 `i`를 사용해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_key().is_equivalent("i")
    `))
})
```

`kwargs`에 할당된 딕셔너리 내포는 `i`의 각 `self`에 대해 `i`의 `other` 속성과 `i`의 `i` 속성의 합을 `vars(self)` 키에 할당해야 합니다. 이를 위해 `getattr()` 함수를 사용하세요.

```js
({ test: () => assert(runPython(`
node = _Node(_code).find_class("R2Vector").find_function("__add__").find_variable("kwargs").find_comp_expr()
node.is_equivalent("getattr(self, i) + getattr(other, i)") or node.is_equivalent("getattr(other, i) + getattr(self, i)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {}
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
