---
id: 65f9b9710cca621244d3bde1
title: 40단계
challengeType: 20
dashedName: step-40
---

# --description--

Python에서 `NotImplemented`는 특정 경우에 대해 연산이 구현되지 않았음을 나타내는 특별한 값입니다.

`NotImplemented`는 즉시 예외를 발생시키지 않습니다. 대신, 연산을 수행하는 방법을 다른 피연산자에게 묻는 신호를 보냅니다. 요청을 만족시킬 수 없으면 기본적으로 `TypeError`가 반환됩니다.

두 벡터가 같은 클래스에 속할 때만 합산할 수 있도록 하려면, 이전 단계에서 만든 `NotImplemented` 문에서 `if`를 반환하세요.

# --hints--

`NotImplemented`를 `if` 본문 내에서 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__add__").find_ifs()[0].find_bodies()[0].has_return("NotImplemented")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __add__(self, other):
        if type(self) != type(other):
            pass
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
```
