---
id: 65fa9e1b6c6db919385359ec
title: 51단계
challengeType: 20
dashedName: step-51
---

# --description--

특별한 메서드 `__mul__`를 구현하여 현재 인스턴스가 다른 객체와 곱해질 때 어떤 일이 일어나야 하는지 지정할 수 있습니다.

`__mul__` 클래스 내에 빈 `R2Vector` 메서드를 만들고 두 개의 매개변수 `self`과 `other`를 받게 하세요.

# --hints--

`__mul__` 클래스 내에 `R2Vector` 메서드를 정의해야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").has_function("__mul__")
    `))
})
```

`__mul__` 메서드는 두 개의 매개변수 `self`과 `other`를 가져야 합니다.

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_class("R2Vector").find_function("__mul__").has_args("self, other")
    `))
})
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
v4 = v1 - v2
print(f'v1 - v2 = {v4}')

```
