---
id: 65fa9402d1fb5516aa42159d
title: 50단계
challengeType: 20
dashedName: step-50
---

# --description--

`v3`를 수정하여 `R2Vector` 인스턴스와 정수나 문자열 같은 다른 유형의 객체를 더해도 됩니다. 콘솔에 `TypeError` 메시지가 출력되는 것을 확인할 수 있습니다.

그런 다음 원래 줄을 복원하고 새 변수 `v4`를 만드세요. `v1`와 `v2`의 차이를 새 변수에 할당하고 이전 f-string과 같은 구조로 결과를 출력하세요.

# --hints--

변수 `v4`를 선언하고 `v1`과 `v2`의 차이를 할당하세요.

```js
({
    test: () => assert(runPython(`
      _Node(_code).has_stmt("v4 = v1 - v2")
    `))
})
```

`f'v1 - v2 = {v4}'`를 출력하세요.

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(f'v1 - v2 = {v4}')")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __add__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) + getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

    def __sub__(self, other):
        if type(self) != type(other):
            return NotImplemented
        kwargs = {i: getattr(self, i) - getattr(other, i) for i in vars(self)}
        return self.__class__(**kwargs)

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z
--fcc-editable-region--
v1 = R2Vector(x=2, y=3)
v2 = R2Vector(x=0.5, y=1.25)
print(f'v1 = {v1}')
print(f'v2 = {v2}')
v3 = v1 + v2
print(f'v1 + v2 = {v3}')
--fcc-editable-region--
```
