---
id: 666827a6fd0dbaafe8330ea6
title: 35단계
challengeType: 20
dashedName: step-35
---

# --description--

`__getattr__`는 속성에 접근하는 데 역할을 하는 또 다른 특별한 메서드입니다.

`__getattribute__`의 기본 구현은 요청한 속성이 인스턴스 속성이 아니거나 클래스 트리에 없을 때 `AttributeError`를 발생시키는 것입니다.

그럴 경우, 클래스에서 정의되어 있다면 `__getattr__`가 호출됩니다. 일반적인 속성 접근 절차가 실패할 때 대체 수단으로 생각할 수 있습니다.

`__getattribute__` 메서드를 `__getattr__`로 바꾸고, 메서드가 반환하는 문자열을 `'calling __getattr__'`로 수정하세요.

# --hints--

`__getattr__` 클래스 안에 `self`라는 이름의 메서드를 두 개의 매개변수 `attr`와 `R2Vector`로 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_args("self, attr")`)) })
```

`__getattr__` 메서드는 문자열 `'calling __getattr__'`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattr__").has_return("'calling __getattr__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattribute__(self, attr):                
        return 'calling __getattribute__'
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

print(v1.x)
print(getattr(v1, 'x'))

```
