---
id: 6668374ed18b7fce10259cb3
title: 37단계
challengeType: 20
dashedName: step-37
---

# --description--

속성 접근 방식에 대해 감을 잡았으니, 코드의 마지막 두 줄과 `__getattr__` 메서드를 함께 삭제하세요.

그다음, 이전에 주석 처리한 `print()` 호출을 복원하고 두 호출 모두에서 두 번째 인수를 삭제하세요.

# --hints--

`__getattr__` 메서드는 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`_Node(_code).find_class("R2Vector").has_function("__getattr__")`)) })
```

코드에 `print(v1.z)`와 `print(getattr(v1, 'z'))`가 없어야 합니다.

```js
({ test: () => assert.isFalse(runPython(`
n = _Node(_code)
n.has_call("print(v1.z)") or n.has_call("print(getattr(v1, 'z'))")
`)) })
```

이전에 주석 처리한 `print()` 호출을 복원하고 두 호출 모두에서 두 번째 인수를 삭제하세요.

```js
({ test: () => assert(runPython(`
_Node(_code).has_call("print(f'v1 = {v1}')") and _Node(_code).has_call("print(f'v2 = {v2}')")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    def __getattr__(self, attr):                
        return 'calling __getattr__'
        
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
print(v1.z)
print(getattr(v1, 'z'))
--fcc-editable-region--
```
