---
id: 66680ddfd0f8c76782923cb0
title: 22단계
challengeType: 20
dashedName: step-22
---

# --description--

`norm()` 메서드는 올바른 값을 반환하고 있지만, 여전히 개선할 수 있는 부분이 있습니다: 가독성입니다.

`vars()` 내장 함수는 객체를 인수로 받아 그 객체의 `__dict__` 속성을 반환합니다.

`__dict__`의 `self` 속성에 직접 접근하는 대신, `norm` 메서드를 수정하여 `vars()` 함수를 사용하세요.

# --hints--

`norm` 메서드를 수정하여 `vars(self)` 대신 `self.__dict__`을 사용하세요.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("norm").has_return("sum(val**2 for val in vars(self).values())**0.5")`)) })
```

# --seed--

## --seed-contents--

```py

class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y
--fcc-editable-region--
    def norm(self):
        return sum(val**2 for val in self.__dict__.values())**0.5
--fcc-editable-region--
    def __str__(self):
        return f'{self.x, self.y}'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
print(v1.norm())
print(v2.norm())

```
