---
id: 666832e427d70bc5219dc62a
title: 36단계
challengeType: 20
dashedName: step-36
---

# --description--

출력 결과에서 볼 수 있듯이, 클래스에서 `__getattr__`를 정의했지만 이 메서드는 아직 호출되지 않았습니다. 이는 기본 속성 접근이 `__getattribute__`를 통해 일어나기 때문입니다. 따라서 터미널에 속성 값이 출력되는 것을 볼 수 있습니다.

이제 코드의 마지막 두 줄을 수정하여 `z` 호출 두 곳 모두에서 `v1`의 `print()` 속성에 접근하세요. 이번에는 `v1`에 없는 속성에 접근할 것입니다. 그 결과 `__getattribute__`가 오류를 발생시키고 `__getattr__`가 호출됩니다.

# --hints--

코드를 수정하여 점 연산자와 `z` 함수를 모두 사용해 `v1`의 `getattr()` 속성을 출력하세요.

```js
({ test: () => assert(runPython(`
n = _Node(_code)
(n.has_call("print(v1.z)") and n.has_call("print(getattr(v1, 'z'))")) or n.has_call("print(v1.z, getattr(v1, 'z'))") or n.has_call("print(getattr(v1, 'z'), v1.z)")
`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'

    def __getattr__(self, attr):                
        return 'calling __getattr__'

class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')
--fcc-editable-region--
print(v1.x)
print(getattr(v1, 'x'))
--fcc-editable-region--
```
