---
id: 66682150151af29efec9727d
title: 33단계
challengeType: 20
dashedName: step-33
---

# --description--

인스턴스 속성에 접근하려 할 때마다 `__getattribute__` 메서드가 내부적으로 호출됩니다. 만약 인스턴스 수준에서 속성을 찾지 못하면, 메서드는 클래스 수준에서, 그리고 결국 부모 클래스들에서 속성을 검색합니다.

클래스 내부에 `__getattribute__` 메서드를 두 개의 매개변수 `self`와 `attr`로 정의하고, 문자열 `'calling __getattribute__'`를 반환하게 만드세요. 속성 접근이 어떻게 작동하는지 보기 위해 기본 구현을 곧 재정의할 것입니다.

# --hints--

`__getattribute__` 클래스 내부에 `self`라는 이름의 메서드를 두 개의 매개변수 `attr`와 `R2Vector`로 정의해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_args("self, attr")`)) })
```

`__getattribute__` 메서드는 문자열 `'calling __getattribute__'`를 반환해야 합니다.

```js
({ test: () => assert(runPython(`_Node(_code).find_class("R2Vector").find_function("__getattribute__").has_return("'calling __getattribute__'")`)) })
```

# --seed--

## --seed-contents--

```py
class R2Vector:
    def __init__(self, *, x, y):
        self.x = x
        self.y = y

    def norm(self):
        return sum(val**2 for val in vars(self).values())**0.5

    def __str__(self):
        return str(tuple(getattr(self, i) for i in vars(self)))

    def __repr__(self):
        arg_list = [f'{key}={val}' for key, val in vars(self).items()]
        args = ', '.join(arg_list)
        return f'{self.__class__.__name__}({args})'
--fcc-editable-region--
    
--fcc-editable-region--
class R3Vector(R2Vector):
    def __init__(self, *, x, y, z):
        super().__init__(x=x, y=y)
        self.z = z

v1 = R2Vector(x=2, y=3)
v2 = R3Vector(x=2, y=2, z=3)
# print(f'v1 = {v1}', f'\nrepr = {repr(v1)}')
# print(f'v2 = {v2}', f'\nrepr = {repr(v2)}')

```
