---
id: 613297a923965e0703b64796
title: Paso 2
challengeType: 0
dashedName: step-2
---

# --description--

Puede que ya esté familiarizado con el elemento `meta`; se utiliza para especificar la información sobre la página, como el título, la descripción, las palabras clave y el autor.

Proporcione a su página un elemento `meta` con un valor de `charset` adecuado.

El atributo `charset` especifica la codificación de caracteres de la página, y, hoy en día, `UTF-8` es la única codificación soportada por la mayoría de los navegadores.

# --hints--

Debe crear un elemento `meta` dentro del elemento `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Debe dar a la etiqueta `meta` un `charset` de `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
