---
id: 6133acc353338c0bba9cb553
title: Paso 5
challengeType: 0
dashedName: step-5
---

# --description--

Finalmente, en el `head`, el elemento `title` es útil para que los lectores de pantalla entiendan el contenido de la página. Además, es una parte importante del *SEO*.

Dale a la página un `title` que sea descriptivo y conciso.

# --hints--

Se debe agregar un elemento `title` al `head`.

```js
assert.exists(document.querySelector('head > title'));
```

No se debe hacer que el `title` tenga más de 60 caracteres.

```js
assert.isAtMost(document.querySelector('head > title')?.textContent?.length, 60);
```

Intenta ser más descriptivo con tu elemento `title`. *Pista: Al menos 15 caracteres*

```js
assert.isAtLeast(document.querySelector('head > title')?.textContent?.length, 15);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <meta name="description" content="freeCodeCamp Accessibility Quiz practice project" />
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
