---
id: 614ccc21ea91ef1736b9b578
title: Paso 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Para este taller, se te ha proporcionado una plantilla básica.

Comienza este viaje de accesibilidad proporcionando un atributo `lang` a tu elemento `html`. Esto ayudará a los lectores de pantalla a identificar el idioma de la página.

También, recuerda enlazar tu hoja de estilos a la página.

# --hints--

You should give the `html` element a `lang` attribute. _Hint: You can use the value `en` for English._

```js
assert.match(code, /<html\s+lang=('|")[\w\-]+?\1\s*>/i);
// TODO: This should/could be fixed in the builder.js
// assert.notThrow(Intl.getCanonicalLocales(document.querySelector('html').lang));
```

Deberías tener un elemento `link` en tu elemento `head`.

```js
assert.exists(document.querySelector("head>link"));
```

Tu elemento `link` debe tener `rel="stylesheet"` y `href="styles.css"`.

```js
const link = document.querySelector("head>link");
assert.equal(link?.getAttribute("rel"), "stylesheet");
const href = link?.getAttribute("data-href");
assert.oneOf(href, ["./styles.css", "styles.css"]);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
--fcc-editable-region--
<html>
  <head>

  </head>
  <body>

  </body>
</html>
--fcc-editable-region--

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
