---
id: 60eebd07ea685b0e777b5583
title: Hatua 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Karibu kwenye mradi wa Fomu ya Usajili! Anza kwa kuongeza tamko la `!DOCTYPE html` juu ya hati ili kivinjari kijue aina ya hati kinachosoma.

# --hints--

Msimbo wako unapaswa kuwa na rejea la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Unapaswa kujumuisha nafasi baada ya rejea la `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Unapaswa kufafanua aina ya hati kuwa `html`.

```js
assert(code.match(/html/gi));
```

Unapaswa kufunga tamko la `DOCTYPE` kwa `>` baada ya aina.

```js
assert(code.match(/html\s*>/gi));
```

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```
