---
id: 60f027099a15b00485563dd2
title: Hatua 2
challengeType: 0
dashedName: step-2
---

# --description--

Chini ya `DOCTYPE`, ongeza kipengele cha `html` chenye sifa ya kipengele `lang` iliyowekwa kuwa `en`, ili uwe na mahali pa kuanza kuweka msimbo.

# --hints--

Tamko lako la `DOCTYPE` linapaswa kuwa mwanzoni mwa HTML yako.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufungua, na sifa ya kipengele `lang` yenye thamani ya `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Kipengele chako cha `html` kinapaswa kuwa na lebo ya kufunga.

```js
assert(code.match(/<\/html\s*>/));
```

Lebos za `html` zinapaswa kuwa katika mpangilio sahihi.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>\s*<\/html\s*>/));
```

Unapaswa kuwa na kipengele kimoja tu cha `html`.

```js
// Possibly a redundant test, as browser fixes this
assert(document.querySelectorAll('html').length === 1);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--
<!DOCTYPE html>

--fcc-editable-region--
```
