---
id: 60f0286404aefb0562a4fdf9
title: Hatua 4
challengeType: 0
dashedName: step-4
---

# --description--

Ongeza kipengele cha `title` na `meta` ndani ya kipengele cha `head`. Toa kichwa kwa mradi wako cha `Registration Form`, na ongeza sifa ya kipengele ya `charset` yenye thamani ya `utf-8` kwa kipengele chako cha `meta`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `title`.

```js
const title = document.querySelector('title');
assert.exists(title);
```

Kipengele chako cha `title` kinapaswa kuwa ndani ya kipengele cha `head`.

```js
assert.exists(document.querySelector('head > title'));
```

Mradi wako unapaswa kuwa na kichwa cha `Registration Form`. Kumbuka, herufi kubwa na ndogo pamoja na tahajia ni muhimu kwa kichwa.

```js
const title = document.querySelector('title');
assert.equal(title.text, 'Registration Form');
```

Unapaswa kuunda kipengele cha `meta` ndani ya kipengele cha `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Unapaswa kutoa kipengele cha `meta` sifa ya kipengele ya `charset` yenye thamani ya `utf-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

Kipengele chako cha `meta` kinapaswa kuwa kipengele kisicho na mwisho, haina lebo ya kumalizia `</meta>`.

```js
assert.notMatch(code, /<\/meta\s*>?/i);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  --fcc-editable-region--
  <head>
    
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
