---
id: 60f02e7361b68405e27b62a5
title: Hatua ya 6
challengeType: 0
dashedName: step-6
---

# --description--

Within the `body`, provide a heading context for the content, by adding an `h1` with the text `Registration Form`.

# --hints--

Unapaswa kuongeza `h1` ndani ya `body`.

```js
assert.exists(document.querySelector('body > h1'));
```

Unapaswa kuipa `h1` maandishi `Registration Form`.

```js
assert.equal(document.querySelector('body > h1')?.textContent, 'Registration Form');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
--fcc-editable-region--
</html>
```
