---
id: 60f1922fcbd2410527b3bd89
title: Hatua 8
challengeType: 0
dashedName: step-8
---

# --description--
Kitengo cha `vh` kinawakilisha urefu wa eneo la kuona, na ni sawa na 1% ya `height` ya eneo la kuona. Hii inafanya iwe kuhusiana na urefu wa eneo la kuona.

Sasa ni wakati wa kupamba mradi kwa kutumia CSS. Anza kwa kumpa `body` `width` ya `100%`, na `height` ya `100vh`.

# --hints--

Unapaswa kutumia kichaguzi cha kipengele cha `body`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('body'));
```

Unapaswa kumpa `body` `width` ya `100%`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.width, '100%');
```

Unapaswa kumpa `body` `height` ya `100vh`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('body')?.height, '100vh');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
  </body>
</html>
```

```css
--fcc-editable-region--

--fcc-editable-region--
```
