---
id: 60f5c3e399ff1a05629964e4
title: Hatua 11
challengeType: 0
dashedName: step-11
---

# --description--

Kama inavyopendekezwa na kichwa, unaunda fomu. Hivyo, baada ya kipengele cha `p`, weka kipengele cha `form` chenye sifa ya kipengele ya `action` inayolenga `https://register-demo.freecodecamp.org`.

# --hints--

Unapaswa kuongeza kipengele cha `form` karibu na kipengele cha `p`.

```js
assert.exists(document.querySelector('p + form'));
```

Unapaswa kumpa `form` sifa ya kipengele ya `action`.

```js
// Default action points to window location
assert.notEqual(document.querySelector('form')?.action, window?.location?.href);
```

Unapaswa kumpa `action` thamani ya `https://register-demo.freecodecamp.org`.

```js
assert.equal(document.querySelector('form')?.action, 'https://register-demo.freecodecamp.org/');
```

Kipengele chako cha `form` kinapaswa kuwa na lebo ya kufunga `</form>`.

```js
assert.match(code, /<\/form\>/);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>

  </body>
--fcc-editable-region--
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}
```
