---
id: 60f803d5241e6a0433a523a1
title: Hatua 16
challengeType: 0
dashedName: step-16
---

# --description--
Kitengo cha `rem` kinawakilisha mzizi `em`, na kinahusiana na ukubwa wa fonti ya kipengele cha `html`.

Kwa kuwa vipengele vya `label` ni ndani ya mstari kwa chaguo-msingi, vyote huonyeshwa pembeni kwa pembeni kwenye mstari mmoja, na kufanya maandishi yao kuwa magumu kusoma. Ili kuonekana kwenye mistari tofauti, ongeza `display: block` kwenye kipengele cha `label`, na ongeza sifa ya `margin` yenye thamani `0.5rem 0`, ili kuvitenganisha.

# --hints--

Unapaswa kutumia kichaguzi cha `label`.

```js
assert.exists(new __helpers.CSSHelp(document).getStyle('label'));
```

Unapaswa kuongeza sifa ya `display` yenye thamani `block`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.display, 'block');
```

Unapaswa kuongeza sifa ya `margin` yenye thamani `0.5rem 0`.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('label')?.margin, '0.5rem 0px');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label>Enter Your First Name:</label>
        <label>Enter Your Last Name:</label>
        <label>Enter Your Email:</label>
        <label>Create a New Password:</label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

--fcc-editable-region--

--fcc-editable-region--
```
