---
id: 60f81616cff80508badf9ad5
title: Hatua 21
challengeType: 0
dashedName: step-21
---

# --description--

Katika hatua hii, unapaswa kuwa na uwezo wa wasilisha fomu. Hata hivyo, unaweza kugundua kuwa hakuna mengi yanayotokea.

Ili kufanya fomu iwe na mwingiliano zaidi, ongeza sifa ya kipengele `required` kwa vipengele vya `input` katika `fieldset` ya kwanza.

Sasa, ikiwa utajaribu wasilisha fomu bila kujaza sehemu zinazohitajika, utaona ujumbe wa kosa.

# --hints--

Unapaswa kumpa kipengele cha kwanza cha `input` sifa ya kipengele `required`.

```js
assert.equal(document.querySelector('input')?.required, true);
```

Unapaswa kumpa kipengele cha pili cha `input` sifa ya kipengele `required`.

```js
assert.equal(document.querySelectorAll('input')?.[1]?.required, true);
```

Unapaswa kumpa kipengele cha tatu cha `input` sifa ya kipengele `required`.

```js
assert.equal(document.querySelectorAll('input')?.[2]?.required, true);
```

Unapaswa kumpa kipengele cha nne cha `input` sifa ya kipengele `required`.

```js
assert.equal(document.querySelectorAll('input')?.[3]?.required, true);
```

Haupasi kumpa `submit` `input` sifa ya kipengele `required`.

```js
assert.equal(document.querySelector('input[type="submit"]')?.required, false);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
--fcc-editable-region--
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" /></label>
      </fieldset>
      <fieldset></fieldset>
      <fieldset></fieldset>
      <input type="submit" value="Submit" />
--fcc-editable-region--
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
