---
id: 62cc5b1779e4d313466f73c5
title: Hatua 5
challengeType: 0
dashedName: step-5
---

# --description--

Panga ndani kipengele cha `link` ndani ya kipengele cha `head`. Mpe sifa ya kipengele ya `rel` yenye thamani ya `stylesheet` na sifa ya kipengele ya `href` yenye thamani ya `styles.css`.

# --hints--

Msimbo wako unapaswa kuwa na kipengele cha `link`.

```js
assert.isNotNull(document.querySelector('link'));
```

Hupaswi kubadilisha lebo zako za `head` zilizopo. Hakikisha hujafuta lebo ya kufunga.

```js
const heads = document.querySelectorAll('head');
assert.strictEqual(heads?.length, 1);
```

Unapaswa kuwa na kipengele kimoja cha `link`.

```js
assert.strictEqual(document.querySelectorAll('link')?.length, 1);
```

Kipengele chako cha `link` kinapaswa kuwa kipengele cha tupu, hakipaswi kuwa na lebo ya kufunga `</link>`.

```js
assert.notMatch(code, /<\/link>/);
```

Kipengele chako cha `link` kinapaswa kuwa ndani ya kipengele chako cha `head`.

```js
const headContentRegex = /(?<=<head\s*>)[\S|\s]*(?=<\/head\s*>)/;
const headElementContent = code.match(headContentRegex);

const headElement = document.createElement("head");
headElement.innerHTML = headElementContent;
assert.isNotNull(headElement.querySelector('link'));
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `rel` yenye thamani `stylesheet`.

```js
const linkElement = document.querySelector('link');
const rel = linkElement?.getAttribute("rel");
assert.strictEqual(rel, "stylesheet");
```

Kipengele chako cha `link` kinapaswa kuwa na sifa ya kipengele ya `href` yenye thamani `styles.css`.

```js
const link = document.querySelector('link');
assert.strictEqual(link?.dataset.href, "styles.css");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
  </head>
--fcc-editable-region--
  <body>
  </body>
</html>
```
