---
id: 62ff8e998d3e7eae14d6ae3b
title: Hatua 30
challengeType: 0
dashedName: step-30
---

# --description--

Unahitaji kuthibitisha kuwa mtumizi amesoma masharti na vigezo.

Ongeza kipengele cha `label`. Ndani ya kipengele kipya cha `label` ongeza kipengele cha `input` na weka sifa ya kipengele ya `type` kuwa `checkbox`. Fanya kipengele hiki cha `input` kuwa `required` ili watumizi wasijisajili bila kukubaliana na masharti na vigezo.

Ongeza `id` na sifa ya kipengele ya `for` yenye thamani ya `terms-and-conditions` kwa vipengele kwa ajili ya ufikikaji.

# --hints--

Unapaswa kuongeza `label` baada ya kipengele cha tatu cha `fieldset`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label'));
```

Unapaswa kuongeza `input` kwa kipengele cha `label`.

```js
assert.exists(document.querySelector('fieldset:nth-child(3) + label > input'));
```

Unapaswa kuongeza sifa ya kipengele ya `type` yenye thamani ya `checkbox` kwa kipengele cha `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.type, 'checkbox');
```

Unapaswa kuongeza sifa ya kipengele ya `required` kwa kipengele cha `input`.

```js
assert.equal(document.querySelector('fieldset:nth-child(3) + label > input')?.required, true);
```

Kipengele cha `input` kinapaswa kuwa na sifa ya kipengele ya `id` yenye thamani ya `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label > input')?.matches('#terms-and-conditions'))
```

Kipengele cha `label` kinapaswa kuwa na sifa ya kipengele ya `for` yenye thamani ya `terms-and-conditions`.

```js
assert(document.querySelector('fieldset:nth-child(3) + label')?.matches('label[for="terms-and-conditions"]'))
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8">
    <title>Registration Form</title>
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <h1>Registration Form</h1>
    <p>Please fill out this form with the required information</p>
    <form method="post" action='https://register-demo.freecodecamp.org'>
      <fieldset>
        <label for="first-name">Enter Your First Name: <input id="first-name" type="text" required /></label>
        <label for="last-name">Enter Your Last Name: <input id="last-name" type="text" required /></label>
        <label for="email">Enter Your Email: <input id="email" type="email" required /></label>
        <label for="new-password">Create a New Password: <input id="new-password" type="password" pattern="[a-z0-5]{8,}" required /></label>
      </fieldset>
      <fieldset>
        <legend>Account type (required)</legend>
        <label for="personal-account"><input id="personal-account" type="radio" name="account-type" checked /> Personal</label>
        <label for="business-account"><input id="business-account" type="radio" name="account-type" /> Business</label>
      </fieldset>
      <fieldset></fieldset>
--fcc-editable-region--

--fcc-editable-region--
      <input type="submit" value="Submit" />
    </form>
  </body>
</html>
```

```css
body {
  width: 100%;
  height: 100vh;
  margin: 0;
  background-color: #1b1b32;
  color: #f5f6f7;
}

label {
  display: block;
  margin: 0.5rem 0;
}

```
