---
id: 686386e76197458b7dfd52ef
title: ステップ 6
challengeType: 0
dashedName: step-6
---

# --description--

これでCTAボタンを作成します。Call to Action（CTA）は、ユーザーに特定の行動を促すためのボタンやリンクで、今回は「聴き始める」ためのものです。

最も内側の`div`の外側に、`href`が`#`に設定されたアンカー要素を作成し、テキストとして`Start listening`を入れてください。

これで最初の料金カードのHTMLは完了です。

# --hints--

アンカー要素を作成してください。

```js
assert.exists(document.querySelector("a"))
```

アンカー要素には`href`が`#`に設定されている必要があります。

```js
const link = document.querySelector("a")
assert.equal(link?.getAttribute("href"), "#")
```

アンカー要素にはテキスト`Start listening`が含まれている必要があります。

```js
assert.equal(document.querySelector("a")?.textContent, "Start listening")
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>
              Start exploring millions of songs with basic features and ads.
            </p>
            <ul>
              <li><span>&#10003;</span>Ad-supported streaming</li>
              <li><span>&#10003;</span>Curated playlists</li>
            </ul>
          </div>
          --fcc-editable-region--
          
          --fcc-editable-region-- 
        </div>
      </div>
    </main>
  </body>
</html>
```
