---
id: 686386e76197458b7dfd52ec
title: ステップ 3
challengeType: 0
dashedName: step-3
---

# --description--

`h2`の下に、テキストが`p`の`$0`要素を作成してください。次に、その`p`要素の中でテキストの後に、テキストが`span`の`/month`要素を作成してください。

# --hints--

テキストが`p`の`$0`要素を作成する必要があります。

```js
const p = document.querySelector("p");
const pTextBeforeSpan = p.childNodes[0].textContent.trim();

assert.equal(pTextBeforeSpan, "$0");
```

`span`要素の中に、テキストが`/month`の`p`要素を作成する必要があります。

```js
const spanText = document.querySelector("p > span")?.textContent;
assert.equal(spanText, "/month");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            --fcc-editable-region--
            
            --fcc-editable-region--
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
