---
id: 686386e76197458b7dfd52eb
title: ステップ 2
challengeType: 0
dashedName: step-2
---

# --description--

`div`の下に`h1`要素を作成してください。その`div`の中にさらに`div`要素を作成し、その中にもう一つ`div`要素をネストしてください。

そして、一番内側の`div`の中に、テキスト`h2`を持つ`Listener`要素を作成してください。これは最初の料金プランを表します。

# --hints--

`div`要素を作成する必要があります。

```js
assert.exists(document.querySelector("div"));
```

外側の`div`の中にもう一つ`div`要素を作成してください。

```js
assert.exists(document.querySelector("div > div"));
```

内側の`div`の中にもう一つ`div`要素を作成してください。

```js
assert.exists(document.querySelector("div > div > div"));
```

一番内側の`h2`要素の中にテキスト`Listener`を持つ`div`を作成してください。

```js
assert.equal(document.querySelector("div > div > div > h2")?.textContent, "Listener");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
    --fcc-editable-region--
      
    --fcc-editable-region--
    </main>
  </body>
</html>
```
