---
id: 686386e76197458b7dfd52ee
title: ステップ 5
challengeType: 0
dashedName: step-5
---

# --description--

`ul` の中に2つのリスト項目を作成してください。最初のリスト項目の中に `span` 要素を作成し、そのテキストとして HTML エンティティ `&#10003;` を追加してください。そのすぐ後に `span` がありますが、リスト項目の中にある状態でテキスト `Ad-supported streaming` を追加してください。

2つ目のリスト項目の中に `span` 要素を作成し、そのテキストとしてエンティティ `&#10003;` を追加してください。そのすぐ後に `span` がありますが、リスト項目の中にある状態でテキスト `Curated playlists` を追加してください。

# --hints--

`ul` 要素の中に2つのリスト項目があるはずです。

```js
assert.lengthOf(document.querySelectorAll("ul > li"), 2)
```

最初の `li` 要素の中に、HTML 文字エンティティ `span` を含む `&#10003;` 要素があるはずです。

```js
const listItems = document.querySelectorAll("ul > li")
const firstSpan = listItems[0].querySelector("span")

assert.equal(firstSpan?.textContent, "✓")
```

最初の `li` 要素の中で、`Ad-supported streaming` 要素のすぐ後にテキスト `span` があるはずです。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[0].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Ad-supported streaming");
```

2つ目の `li` 要素の中に、HTML 文字エンティティ `span` を含む `&#10003;` 要素があるはずです。

```js
const listItems = document.querySelectorAll("ul > li")
const secondSpan = listItems[1].querySelector("span")

assert.equal(secondSpan?.textContent, "✓")
```

2つ目の `li` 要素の中で、`Curated playlists` 要素のすぐ後にテキスト `span` があるはずです。

```js
const listItems = document.querySelectorAll("ul > li");
const span = listItems[1].querySelector("span");

assert.equal(span.nextSibling?.textContent.trim(), "Curated playlists");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
      <h1>Choose your listening plan</h1>
      <div>
        <div>
          <div>
            <h2>Listener</h2>
            <p>$0<span>/month</span></p>
            <p>Start exploring millions of songs with basic features and ads.</p>
            <ul>
            --fcc-editable-region--
              
            --fcc-editable-region--
            </ul>
          </div>
        </div>
      </div>
    </main>
  </body>
</html>
```
