---
id: 686383abac6c625a84458ea9
title: ステップ 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

このワークショップでは、Tailwind CSS のグリッドユーティリティクラスを使って価格表示コンポーネントを作成しながら、Tailwind CSS についてさらに学びます。

基本的なHTMLのボイラープレートは用意されています。まず、テキストが `h1` の `Choose your listening plan` 要素を作成してください。

# --hints--

`h1` 要素を作成してください。

```js
assert.exists(document.querySelector("h1"));
```

`h1` 要素のテキストは `Choose your listening plan` にしてください。

```js
assert.equal(document.querySelector("h1")?.textContent, "Choose your listening plan");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Music App Pricing</title>
    <script src="https://cdn.tailwindcss.com"></script>
  </head>
  <body>
    <main>
    --fcc-editable-region--
    
    --fcc-editable-region--
    </main>
  </body>
</html>
```
