---
id: 587d78ad367417b2b2512afa
title: Tumia sifa ya flex-wrap kufunika safu au nguzo
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cQv9ZtG'
forumTopicId: 301114
dashedName: use-the-flex-wrap-property-to-wrap-a-row-or-column
---

# --description--

CSS flexbox ina kipengele cha kugawanya kontena cha kubadilika katika safu nyingi (au nguzo). Kwa chaguo-msingi, kontena cha kubadilika kitaweka vipande vyote vya kubadilika pamoja. Kwa mfano, safu yote itakuwa kwenye mstari mmoja.

Hata hivyo, kutumia sifa ya `flex-wrap` hueleza CSS kufunika vipande. Hii inamaanisha vipande vya ziada vinahamia kwenye safu au nguzo mpya. Hatua ya kusimama ambapo kufunika kunatokea hutegemea ukubwa wa vipande na ukubwa wa kontena.

CSS pia ina chaguzi za mwelekeo wa kufunika:

<ul><li><code>nowrap</code>: hii ni mipangilio ya chaguo-msingi, na haifuniki vipande.</li><li><code>wrap</code>: hufunika vipande kwenye mistari mingi kutoka juu hadi chini ikiwa ni safu na kutoka kushoto kwenda kulia ikiwa ni nguzo.</li><li><code>wrap-reverse</code>: hufunika vipande kwenye mistari mingi kutoka chini hadi juu ikiwa ni safu na kutoka kulia kwenda kushoto ikiwa ni nguzo.</li></ul>

# --instructions--

Mpangilio wa sasa una masanduku mengi sana kwa safu moja. Ongeza sifa ya CSS `flex-wrap` kwenye kipengele cha `#box-container`, na uipe thamani ya `wrap`.

# --hints--

Kipengele cha `#box-container` kinapaswa kuwa na sifa ya `flex-wrap` ikiwa na thamani ya `wrap`.

```js
const boxContainer = document.querySelector('#box-container');
const flexWrap = window.getComputedStyle(boxContainer)['flex-wrap'];
assert.strictEqual(flexWrap, 'wrap');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    background: gray;
    display: flex;
    height: 100%;
    flex-wrap: wrap;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 25%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 25%;
    height: 50%;
  }
  #box-3 {
    background-color: violet;
    width: 25%;
    height: 50%;
  }
  #box-4 {
    background-color: yellow;
    width: 25%;
    height: 50%;
  }
  #box-5 {
    background-color: green;
    width: 25%;
    height: 50%;
  }
  #box-6 {
    background-color: black;
    width: 25%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
  <div id="box-3"></div>
  <div id="box-4"></div>
  <div id="box-5"></div>
  <div id="box-6"></div>
</div>
```
