---
id: 587d78ae367417b2b2512afd
title: Tumia sifa ya flex-basis kuweka ukubwa wa awali wa kipengee
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c3d9nCa'
forumTopicId: 301108
dashedName: use-the-flex-basis-property-to-set-the-initial-size-of-an-item
---

# --description--

Sifa ya `flex-basis` inaeleza ukubwa wa awali wa kipengee kabla ya CSS kufanya marekebisho kwa `flex-shrink` au `flex-grow`.

Vipimo vinavyotumika na sifa ya `flex-basis` ni sawa na sifa nyingine za ukubwa (`px`, `em`, `%`, n.k.). Thamani ya `auto` huweka ukubwa wa vipengee kulingana na yaliyomo.

# --instructions--

Weka ukubwa wa awali wa masanduku kwa kutumia `flex-basis`. Ongeza sifa ya CSS `flex-basis` kwa `#box-1` na `#box-2`. Toa `#box-1` thamani ya `10em` na `#box-2` thamani ya `20em`.

# --hints--

Kipengee cha `#box-1` kinapaswa kuwa na sifa ya `flex-basis`.

```js
const boxOne = document.querySelector('#box-1');
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Kipengee cha `#box-1` kinapaswa kuwa na thamani ya `flex-basis` ya `10em`.

```js
assert.match(code, /#box-1\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?10em;/g);
```

Kipengee cha `#box-2` kinapaswa kuwa na sifa ya `flex-basis`.

```js
const boxTwo = document.querySelector('#box-2');
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];
assert.notStrictEqual(flexBasis, 'auto');
```

Kipengee cha `#box-2` kinapaswa kuwa na thamani ya `flex-basis` ya `20em`.

```js
assert.match(code, /#box-2\s*?{\s*?.*?\s*?.*?\s*?flex-basis:\s*?20em;/g);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-basis: 10em;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-basis: 20em;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
