---
id: 587d78ac367417b2b2512af4
title: Tumia sifa ya flex-direction kutengeneza safu wima
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cZmWeA4'
forumTopicId: 301109
dashedName: use-the-flex-direction-property-to-make-a-column
---

# --description--

Changamoto mbili za mwisho zilitumia seti ya sifa `flex-direction` iliyowekwa kuwa `row`. Sifa hii pia inaweza kutengeneza safu kwa kupanga watoto wa kontena cha kubadilika wima.

# --instructions--

Ongeza sifa ya CSS `flex-direction` kwenye kipengele cha `#box-container`, na uipe thamani ya `column`.

# --hints--

Kipengele cha `#box-container` kinapaswa kuwa na sifa ya `flex-direction` iliyowekwa kuwa `column`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'column');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: column;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
