---
id: 587d78ab367417b2b2512af0
title: 'Tumia display: flex kuweka sanduku mbili'
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cgz3QS7'
forumTopicId: 301105
dashedName: use-display-flex-to-position-two-boxes
---

# --description--

Sehemu hii inatumia mitindo mbadala ya changamoto kuonyesha jinsi ya kutumia CSS kuweka vipengele kwa njia inayojibadilisha kulingana na kifaa. Kwanza, changamoto itafafanua nadharia, kisha changamoto ya vitendo ikitumia sehemu rahisi ya tweet itatumia dhana ya flexbox.

Kuweka sifa ya CSS `display: flex;` kwenye kipengele kunakuwezesha kutumia sifa nyingine za flex kujenga ukurasa unaojibadilisha kulingana na kifaa.

# --instructions--

Ongeza sifa ya CSS `display` kwenye `#box-container` na weka thamani yake kuwa `flex`.

# --hints--

`#box-container` inapaswa kuwa na sifa ya `display` ikiwa na thamani ya `flex`.

```js
const boxContainer = document.querySelector('#box-container');
const displayStyle = window.getComputedStyle(boxContainer)['display'];
assert.strictEqual(displayStyle, 'flex');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    height: 500px;

  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    height: 500px;
    display: flex;
  }

  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>
<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
