---
id: 587d78ab367417b2b2512af2
title: Tumia sifa ya flex-direction kutengeneza safu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cBEkbfJ'
forumTopicId: 301110
dashedName: use-the-flex-direction-property-to-make-a-row
---

# --description--

Kuongeza `display: flex` kwenye kipengele hubadilisha kuwa chombo cha kubadilika. Hii inafanya iwezekane kupanga watoto wowote wa kipengele hicho katika safu au nguzo. Unafanya hivi kwa kuongeza sifa ya `flex-direction` kwenye kipengele mzazi na kuiweka kuwa row au column. Kutengeneza safu kutapanga watoto kwa usawa, na kutengeneza nguzo kutapanga watoto kwa wima.

Chaguzi nyingine za `flex-direction` ni `row-reverse` na `column-reverse`.

**Kumbuka:** Thamani ya msingi ya sifa ya `flex-direction` ni `row`.

# --instructions--

Ongeza sifa ya CSS `flex-direction` kwenye kipengele cha `#box-container`, na uipe thamani ya `row-reverse`.

# --hints--

Kipengele cha `#box-container` kinapaswa kuwa na sifa ya `flex-direction` iliyowekwa kuwa `row-reverse`.

```js
const boxContainer = document.querySelector('#box-container');
const flexDirection = window.getComputedStyle(boxContainer)['flex-direction'];
assert.strictEqual(flexDirection, 'row-reverse');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;

  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
    flex-direction: row-reverse;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 50%;
    height: 50%;
  }

  #box-2 {
    background-color: orangered;
    width: 50%;
    height: 50%;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
