---
id: 587d78ae367417b2b2512afe
title: Tumia sifa ya ufupisho ya flex
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cbpW2tE'
forumTopicId: 301112
dashedName: use-the-flex-shorthand-property
---

# --description--

Kuna njia ya mkato inayopatikana kuweka sifa kadhaa za flex kwa pamoja. Sifa za `flex-grow`, `flex-shrink`, na `flex-basis` zote zinaweza kuwekwa kwa pamoja kwa kutumia sifa ya `flex`.

Kwa mfano, `flex: 1 0 10px;` itaweka kipengee kuwa na `flex-grow: 1;`, `flex-shrink: 0;`, na `flex-basis: 10px;`.

Mipangilio ya sifa kwa kawaida ni `flex: 0 1 auto;`.

# --instructions--

Ongeza sifa ya CSS `flex` kwa `#box-1` na `#box-2`. Mpe `#box-1` thamani ili `flex-grow` yake iwe `2`, `flex-shrink` yake iwe `2`, na `flex-basis` yake iwe `150px`. Mpe `#box-2` thamani ili `flex-grow` yake iwe `1`, `flex-shrink` yake iwe `1`, na `flex-basis` yake iwe `150px`.

Thamani hizi zitasababisha `#box-1` kukua kujaza nafasi ya ziada kwa mara mbili ya `#box-2` wakati kontena ni kubwa zaidi ya 300px na kupungua kwa mara mbili ya `#box-2` wakati kontena ni chini ya 300px. 300px ni ukubwa wa pamoja wa thamani za `flex-basis` za maboksi mawili.

# --hints--

Kipengee cha `#box-1` kinapaswa kuwa na sifa ya `flex` iliyowekwa kwa thamani ya `2 2 150px`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxOne)['flex-basis'];

assert.equal(flexGrow, '2');
assert.equal(flexShrink, '2');
assert.equal(flexBasis, '150px');
```

Kipengee cha `#box-2` kinapaswa kuwa na sifa ya `flex` iliyowekwa kwa thamani ya `1 1 150px`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
const flexBasis = window.getComputedStyle(boxTwo)['flex-basis'];

assert.equal(flexGrow, '1');
assert.equal(flexShrink, '1');
assert.equal(flexBasis, '150px');
```

Msimbo wako unapaswa kutumia sifa ya `flex` kwa `#box-1` na `#box-2`.

```js
assert.lengthOf(code.match(/flex:\s*?\d\s+?\d\s+?150px;/g), 2);
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    flex: 2 2 150px;
    height: 200px;
  }

  #box-2 {
    background-color: orangered;
    flex: 1 1 150px;
    height: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
