---
id: 587d78af367417b2b2512b00
title: Tumia sifa ya align-self
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvzfv'
forumTopicId: 301107
dashedName: use-the-align-self-property
---

# --description--

Sifa ya mwisho kwa vipande vya kubadilika ni `align-self`. Sifa hii inakuwezesha kurekebisha ulinganifu wa kila kipande binafsi, badala ya kuziweka zote kwa pamoja. Hii ni muhimu kwa sababu mbinu nyingine za kawaida za marekebisho zinazotumia sifa za CSS `float`, `clear`, na `vertical-align` hazifanyi kazi kwa vipande vya kubadilika.

`align-self` inakubali thamani sawa na `align-items` na itaghairi thamani yoyote iliyowekwa na sifa ya `align-items`.

# --instructions--

Ongeza sifa ya CSS `align-self` kwa `#box-1` na `#box-2`. Mpe `#box-1` thamani ya `center` na mpe `#box-2` thamani ya `flex-end`.

# --hints--

Kipengele cha `#box-1` kinapaswa kuwa na sifa ya `align-self` ikiwa na thamani ya `center`.

```js
const boxOne = document.querySelector('#box-1');
const alignment = window.getComputedStyle(boxOne)['align-self'];
assert.strictEqual(alignment, 'center');
```

Kipengele cha `#box-2` kinapaswa kuwa na sifa ya `align-self` ikiwa na thamani ya `flex-end`.

```js
const boxTwo = document.querySelector('#box-2');
const alignment = window.getComputedStyle(boxTwo)['align-self'];
assert.strictEqual(alignment, 'flex-end');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    align-self: center;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    align-self: flex-end;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
