---
id: 587d78ad367417b2b2512afb
title: Tumia sifa ya flex-shrink kupunguza ukubwa wa vipande
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cd3PBfr'
forumTopicId: 301113
dashedName: use-the-flex-shrink-property-to-shrink-items
---

# --description--

Hadi sasa, sifa zote katika changamoto zinahusu chombo cha kubadilika (mzazi wa vipande vya kubadilika). Hata hivyo, kuna sifa kadhaa muhimu kwa vipande vya kubadilika.

Ya kwanza ni sifa ya `flex-shrink`. Inapotumika, inaruhusu kipande kupungua ukubwa ikiwa chombo cha kubadilika ni kidogo sana. Vipande hupungua ukubwa wakati upana wa kontena mzazi ni mdogo kuliko jumla ya upana wa vipande vyote vya kubadilika ndani yake.

Sifa ya `flex-shrink` huchukua nambari kama thamani. Nambari kubwa zaidi, ndivyo kipande kitakapopungua ukubwa zaidi ikilinganishwa na vipande vingine ndani ya kontena. Kwa mfano, ikiwa kipande kimoja kina thamani ya `flex-shrink` ya `1` na kingine kina thamani ya `flex-shrink` ya `3`, kile chenye thamani ya `3` kitapungua ukubwa mara tatu zaidi kuliko kingine.

# --instructions--

Ongeza sifa ya CSS `flex-shrink` kwa `#box-1` na `#box-2`. Wape `#box-1` thamani ya `1` na `#box-2` thamani ya `2`.

# --hints--

Kipengele cha `#box-1` kinapaswa kuwa na sifa ya `flex-shrink` ikiwa na thamani ya `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexShrink = window.getComputedStyle(boxOne)['flex-shrink'];
assert.equal(flexShrink, '1');
```

Kipengele cha `#box-2` kinapaswa kuwa na sifa ya `flex-shrink` ikiwa na thamani ya `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexShrink = window.getComputedStyle(boxTwo)['flex-shrink'];
assert.equal(flexShrink, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    width: 100%;
    height: 200px;
    flex-shrink: 1;
  }

  #box-2 {
    background-color: orangered;
    width: 100%;
    height: 200px;
    flex-shrink: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
