---
id: 587d78ae367417b2b2512aff
title: Tumia sifa ya order kupanga upya vipande
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/cMbvNAG'
forumTopicId: 301116
dashedName: use-the-order-property-to-rearrange-items
---

# --description--

Sifa ya `order` hutumika kuambia CSS mpangilio wa jinsi vipande vya kubadilika vinaonekana katika chombo cha kubadilika. Kwa chaguo-msingi, vipande vitaonekana kwa mpangilio ule ule walipotokea katika chanzo cha HTML. Sifa hii hupokea nambari kama thamani, na nambari hasi zinaweza kutumika.

# --instructions--

Ongeza sifa ya CSS `order` kwa `#box-1` na `#box-2`. Toa `#box-1` thamani ya `2` na `#box-2` thamani ya `1`.

# --hints--

Kipande cha kubadilika `#box-1` kinapaswa kuwa na sifa ya `order` ikiwa na thamani ya `2`.

```js
const boxOne = document.querySelector('#box-1');
const order = window.getComputedStyle(boxOne)['order'];
assert.strictEqual(order, '2');
```

Kipande cha kubadilika `#box-2` kinapaswa kuwa na sifa ya `order` ikiwa na thamani ya `1`.

```js
const boxTwo = document.querySelector('#box-2');
const order = window.getComputedStyle(boxTwo)['order'];
assert.strictEqual(order, '1');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;

    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;

    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }
  #box-1 {
    background-color: dodgerblue;
    order: 2;
    height: 200px;
    width: 200px;
  }

  #box-2 {
    background-color: orangered;
    order: 1;
    height: 200px;
    width: 200px;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
