---
id: 587d78ae367417b2b2512afc
title: Tumia sifa ya flex-grow kupanua vitu
challengeType: 0
videoUrl: 'https://scrimba.com/p/pVaDAv/c2p78cg'
forumTopicId: 301111
dashedName: use-the-flex-grow-property-to-expand-items
---

# --description--

Kinyume cha `flex-shrink` ni sifa ya `flex-grow`. Kumbuka kwamba `flex-shrink` hudhibiti ukubwa wa vitu wakati kontena linapokaza. Sifa ya `flex-grow` hudhibiti ukubwa wa vitu wakati kontena mzazi unapopanuka.

Kutumia mfano kama ule wa changamoto iliyopita, ikiwa kipengee kimoja kina thamani ya `flex-grow` ya `1` na kingine kina thamani ya `flex-grow` ya `3`, kile chenye thamani ya `3` kitakua mara tatu zaidi kuliko kingine.

# --instructions--

Ongeza sifa ya CSS `flex-grow` kwa `#box-1` na `#box-2`. Toa `#box-1` thamani ya `1` na `#box-2` thamani ya `2`.

# --hints--

Kipengee cha `#box-1` kinapaswa kuwa na sifa ya `flex-grow` ikiwa na thamani ya `1`.

```js
const boxOne = document.querySelector('#box-1');
const flexGrow = window.getComputedStyle(boxOne)['flex-grow'];
assert.equal(flexGrow, '1');
```

Kipengee cha `#box-2` kinapaswa kuwa na sifa ya `flex-grow` ikiwa na thamani ya `2`.

```js
const boxTwo = document.querySelector('#box-2');
const flexGrow = window.getComputedStyle(boxTwo)['flex-grow'];
assert.equal(flexGrow, '2');
```

# --seed--

## --seed-contents--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;

  }

  #box-2 {
    background-color: orangered;
    height: 200px;

  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```

# --solutions--

```html
<style>
  #box-container {
    display: flex;
    height: 500px;
  }

  #box-1 {
    background-color: dodgerblue;
    height: 200px;
    flex-grow: 1;
  }

  #box-2 {
    background-color: orangered;
    height: 200px;
    flex-grow: 2;
  }
</style>

<div id="box-container">
  <div id="box-1"></div>
  <div id="box-2"></div>
</div>
```
