---
id: 66c06d618d075c7f7f1b890a
title: Crea un indovina fortuna
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

In questo laboratorio, definirai cinque fortune e ne selezionerai una a caso da mostrare all'utente.

**Obiettivo:** Completa le storie utente qui sotto e fai passare tutti i test per completare il laboratorio.

**Storie utente:**

1. Dovresti inizializzare le cinque variabili `fortune1`, `fortune2`, `fortune3`, `fortune4` e `fortune5` con una stringa a tua scelta. Puoi usare le opzioni qui sotto se vuoi:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Dovresti selezionare un numero casuale tra 1 e 5, inclusi, e assegnarlo alla variabile `randomNumber`.

3. Dovresti creare una variabile `selectedFortune` e assegnare la fortuna appropriata basandoti su queste regole:

   - Se `randomNumber` è 1, assegna il valore di `fortune1` a `selectedFortune`.
   - Se `randomNumber` è 2, assegna il valore di `fortune2` a `selectedFortune`.
   - Se `randomNumber` è 3, assegna il valore di `fortune3` a `selectedFortune`.
   - Se `randomNumber` è 4, assegna il valore di `fortune4` a `selectedFortune`.
   - Se `randomNumber` è 5, assegna il valore di `fortune5` a `selectedFortune`.

4. Dovresti stampare `selectedFortune` nella console.

# --hints--

Dovresti inizializzare `fortune1` con una stringa.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Dovresti inizializzare `fortune2` con una stringa.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Dovresti inizializzare `fortune3` con una stringa.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Dovresti inizializzare `fortune4` con una stringa.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Dovresti inizializzare `fortune5` con una stringa.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Dovresti usare il metodo `Math.random()` per generare un numero casuale.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Dovresti generare un numero casuale tra 1 e 5, inclusi, e assegnarlo alla variabile `randomNumber`.

```js
assert.isNotNull(randomNumber);
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length]
  const userCode__test_code__ = new Function(`${code};return randomNumber`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__()
    assert.strictEqual(result, i + 1);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Dovresti avere una variabile `selectedFortune` a cui viene assegnato un valore basato sul valore di `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

`randomNumber` dovrebbe corrispondere alla sua fortuna. Per esempio, se `randomNumber` è 1, `selectedFortune` dovrebbe essere uguale a `fortune1` e così via.

```js
const originalRandom__test_code__ = Math.random;
try {
  const mockValues__test_code__ = [0.0, 0.2, 0.4, 0.6, 0.8];
  let callCount__test_code__ = 0;
  Math.random = () => mockValues__test_code__[callCount__test_code__++ % mockValues__test_code__.length];
  const fortunes__test_code__ = [fortune1, fortune2, fortune3, fortune4, fortune5,];
  const userCode__test_code__ = new Function(`${code};return selectedFortune`);
  for (let i = 0; i < 5; i++) {
    const result = userCode__test_code__();
    assert.strictEqual(result, fortunes__test_code__[i]);
  }
} finally {
  Math.random = originalRandom__test_code__;
}
```

Dovresti stampare `selectedFortune` nella console.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
