---
id: 694a9ff65d4472dd451d8255
title: 1단계
challengeType: 20
dashedName: step-1
---

# --description--

이 워크숍에서는 영화 티켓 예약 계산기를 만들면서 Python에서 불리언과 조건문을 사용하는 방법을 연습합니다.

사용자와 영화 상영에 관한 정보를 저장하는 변수를 선언하는 것부터 시작하세요.

이전 수업에서 배운 것처럼, 변수는 대입 연산자(`=`)를 사용하여 다음과 같이 할당합니다:

```py
# Assigns the value 10 to variable x
x = 10
```

영화 티켓의 기본 가격을 저장할 `base_price`라는 변수를 만들고 그 값을 `15`로 설정하세요. 또 다른 변수 `age`를 만들어 사용자의 나이를 저장하고 그 값을 `21`로 설정하세요.

# --hints--

`base_price`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("base_price")`)) })
```

정수 값 `15`를 `base_price`에 할당해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("base_price").is_equivalent("base_price = 15")
`)) })
```

`age`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("age")`)) })
```

정수 값 `21`를 `age`에 할당해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("age").is_equivalent("age = 21")
`)) })
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
