---
id: 694cc1b515976889afd4cded
title: 2단계
challengeType: 20
dashedName: step-2
---

# --description--

이제 영화 티켓에 관한 문자열 값을 저장해야 합니다. Python에서는 다음과 같이 문자열을 저장할 수 있습니다:

```py
name = 'Alice'
```

사용자가 선택한 좌석 유형을 저장할 변수 이름을 `seat_type`로 만들고, 그 값을 문자열 `Gold`로 설정하세요. 영화 상영 시간을 저장할 또 다른 변수 이름을 `show_time`로 만들고, 그 값을 문자열 `Evening`로 설정하세요.

두 값 모두를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

# --hints--

`seat_type`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("seat_type")`)) })
```

문자열 `Gold`를 `seat_type`에 할당해야 합니다. 값은 작은따옴표나 큰따옴표로 감싸야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("seat_type").is_equivalent("seat_type = 'Gold'")
`)) })
```

`show_time`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("show_time")`)) })
```

문자열 `Evening`를 `show_time`에 할당해야 합니다. 값은 작은따옴표나 큰따옴표로 감싸야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("show_time").is_equivalent("show_time = 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
--fcc-editable-region--

--fcc-editable-region--
```
