---
id: 6958ea03cb1c4684c138754a
title: 21단계
challengeType: 20
dashedName: step-21
---

# --description--

마지막 단계에서는 이전 단계에서 계산한 값을 사용해 영화 티켓의 최종 가격을 계산할 것입니다.

최종 티켓 가격은 기본 가격에 추가 요금과 서비스 요금을 더한 후 할인을 빼서 계산합니다.

마지막 `if` 문장 본문 안에서 `print('Service charges:', service_charges)` 줄 아래에 티켓의 최종 가격을 계산해 `final_price`라는 변수에 저장하세요.

마지막으로 `Final price of ticket:` 다음에 `final_price` 값을 출력하는 메시지를 출력하세요.

# --hints--

`final_price`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("final_price")`)) })
```

`base_price + extra_charges + service_charges - discount`에 `final_price` 값을 할당해야 합니다.

```js
({ test: () => (runPython(`
import itertools  
perms = itertools.permutations(['+ base_price', '+ extra_charges', '+ service_charges', '- discount'])  
values = (' '.join(perm).lstrip('+') for perm in perms)  
solutions = (f'final_price = {v}' for v in values)  
var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("final_price")  
assert any(var.is_equivalent(s) for s in solutions)
`)) })
```

`print('Final price of ticket:', final_price)` 변수 아래에 `final_price`가 있어야 합니다.

```js
({ test: () => (runPython(`
print_call = "print('Final price of ticket:', final_price)"  
_body = _Node(_code).find_ifs()[4].find_bodies()[0]
assert _body.has_call(print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

    service_charges = 0
    if seat_type == 'Premium':
        service_charges = 5
    elif seat_type == 'Gold':
        service_charges = 3
    else:
        service_charges = 1
    print('Service charges:', service_charges)

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```

# --solutions--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

    service_charges = 0
    if seat_type == 'Premium':
        service_charges = 5
    elif seat_type == 'Gold':
        service_charges = 3
    else:
        service_charges = 1
    print('Service charges:', service_charges)

    final_price = base_price + extra_charges + service_charges - discount
    print('Final price of ticket:', final_price)
else:
    print('Ticket booking failed due to restrictions')
```
