---
id: 6958dfa0c74b237e1b56ce04
title: 7단계
challengeType: 20
dashedName: step-7
---

# --description--

사용자가 회원이면 회원 할인 혜택을 받을 수 있습니다.

`discount`라는 변수를 만들고 그 값을 `0`로 설정하세요. 이 변수는 사용자가 영화 티켓에 받는 할인을 저장합니다.

# --hints--

`discount`라는 변수를 선언하세요.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("discount")`)) })
```

정수 값 `0`를 `discount`에 할당하세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("discount").is_equivalent("discount = 0")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

--fcc-editable-region--

--fcc-editable-region--
```
