---
id: 6958dd85587e10c99162dd81
title: 3단계
challengeType: 20
dashedName: step-3
---

# --description--

이전 수업에서 배운 것처럼, Python에서는 조건이 참일 때 코드를 실행하기 위해 `if` 문을 사용할 수 있습니다.

`if` 문은 `if` 키워드 다음에 조건과 콜론이 옵니다. 조건이 참일 때 실행할 코드는 들여쓰기 되어야 하며, 이를 `if` 문의 본문이라고 합니다.

```py
if condition:
    # Code to execute if condition is True
```

이번 단계에서는 사용자의 나이를 기준으로 영화 티켓을 예약할 수 있는지 확인할 것입니다.

`if` 문을 만들어 `age` 가 `17` 보다 큰지 확인하세요. `if` 문의 본문 안에서 `User is eligible to book a ticket` 을 출력하세요. 이렇게 하면 사용자의 나이가 `17` 보다 클 때만 메시지가 출력됩니다.

`if` 문의 본문을 반드시 들여쓰기하고, `print()` 호출 안에서 메시지를 작은따옴표나 큰따옴표로 감싸세요.

# --hints--

`if` 문이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[0]`)) })
```

`if` 문의 조건은 `age > 17` 여야 합니다.

```js 
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("age > 17") or _cond.is_equivalent("17 < age")
`)) })
```

`print()` 문 안에 `if` 호출이 있어야 합니다. `if` 문의 본문을 들여쓰기해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].find_calls("print")[0]
`)) })
```

`User is eligible to book a ticket` 문 안에서 `if` 을 출력해야 합니다. 메시지는 작은따옴표나 큰따옴표로 감싸세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[0].find_bodies()[0].has_call("print('User is eligible to book a ticket')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

--fcc-editable-region--

--fcc-editable-region--
```
