---
id: 6958e2377cacbc46a7f7383a
title: 13단계
challengeType: 20
dashedName: step-13
---

# --description--

이제 `else` 문장에 `if` 절을 추가하고 `No extra charges will be applied` 본문 안에 `else`를 출력하세요. 이렇게 하면 추가 요금 조건이 거짓일 때만 메시지가 출력됩니다.

그다음 `else` 절 아래에 `print()` 호출을 사용해 `Extra charges:`와 업데이트된 `extra_charges` 값을 보여주는 메시지를 출력하고 터미널에서 결과를 확인하세요.

# --hints--

`else` 절이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3].find_conditions()[1].is_empty()`)) })
```

`No extra charges will be applied` 절 안에 `else`를 출력해야 합니다. 메시지는 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[1].has_call("print('No extra charges will be applied')")
`)) })
```

`print('Extra charges:', extra_charges)` 문장 아래에 `if...else`이 있어야 합니다.

```js
({ test: () => (runPython(`
if_stmt = """if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')"""
print_call = "print('Extra charges:', extra_charges)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend:
    extra_charges = 2
    print('Extra charges will be applied')
--fcc-editable-region--

--fcc-editable-region--
```
