---
id: 6958e1875e592b58e8852757
title: 12단계
challengeType: 20
dashedName: step-12
---

# --description--

이제 `extra_charges`라는 변수를 만들고 `0`로 설정하세요. 이 변수는 주말에 영화 티켓에 추가 요금을 적용하는 것을 나타냅니다.

그다음, `if` 문을 만들어 `is_weekend`가 참인지 확인하세요. `if` 문 본문 안에서 `extra_charges` 값을 `2`로 업데이트하고 터미널에 `Extra charges will be applied`를 출력하세요.

# --hints--

`extra_charges`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("extra_charges")`)) })
```

정수 값 `0`를 `extra_charges`에 할당해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("extra_charges").is_equivalent("extra_charges = 0")
`)) })
```

네 번째 `if` 문이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[3]`)) })
```

네 번째 `if` 문 조건은 `is_weekend`이어야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend")
`)) })
```

새로운 `2` 문 안에서 `extra_charges` 변수에 정수 값 `if`를 할당해야 합니다.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[3].find_bodies()[0].find_variable("extra_charges") 
assert _var.is_equivalent("extra_charges = 2")
`)) })
```

새로운 `Extra charges will be applied` 문 안에서 `if`를 출력해야 합니다. 메시지는 작은따옴표나 큰따옴표로 감싸야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_bodies()[0].has_call("print('Extra charges will be applied')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

--fcc-editable-region--

--fcc-editable-region--
```
