---
id: 6958e0ecc121b0ae3dd9c5e0
title: 9단계
challengeType: 20
dashedName: step-9
---

# --description--

사용자가 멤버십 할인 대상이 아닐 때의 경우도 처리해야 합니다.

`else` 문장에 `if` 절을 추가하고 `User does not qualify for membership discount` 본문 안에 `else`를 출력하세요.

변수 뒤에 메시지를 출력하려면 이렇게 하세요:

```py
print('Total:', total)
```

또한 `discount`의 업데이트된 값을 표시해야 합니다. `if...else` 문장 아래에 `print()` 호출을 사용해 `Discount:` 메시지와 `discount`의 업데이트된 값을 출력하세요. 그런 다음 터미널에서 출력 결과를 확인하세요.

# --hints--

`else` 절이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2].find_conditions()[1].is_empty()`)) })
```

`User does not qualify for membership discount` 절 안에 `else`을 출력해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[1].has_call("print('User does not qualify for membership discount')")
`)) })
```

`print('Discount:', discount)` 문장 아래에 `if...else`가 있어야 합니다.

```js
({ test: () => (runPython(`
if_stmt = """if is_member:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')"""
print_call = "print('Discount:', discount)"
assert _Node(_code).is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
if is_member:
    discount = 3
    print('User qualifies for membership discount')
--fcc-editable-region--

--fcc-editable-region--
```
