---
id: 6958de0d773289a89515a6b8
title: 4단계
challengeType: 20
dashedName: step-4
---

# --description--

이제 사용자의 나이를 기준으로 저녁 공연 예약이 가능한지 확인하세요.

`if` 문을 만들어 `age`이 `21`보다 크거나 같은지 확인하세요. `if` 문 안에서는 `User is eligible for Evening shows`를 출력하세요.

# --hints--

코드에 `if` 문이 두 개 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1]`)) })
```

두 번째 `if` 문 조건은 `age >= 21`이어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[1].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

새로운 `User is eligible for Evening shows` 문 안에서 `if`를 출력하세요. 메시지는 작은따옴표나 큰따옴표로 감싸야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[0].has_call("print('User is eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

--fcc-editable-region--

--fcc-editable-region--
```
