---
id: 6958e59845777391b86d7331
title: 18단계
challengeType: 20
dashedName: step-18
---

# --description--

Python에서 `if` 문은 다른 `if` 문 본문 안에 배치할 수도 있습니다. 이를 중첩된 `if` 문이라고 합니다.

중첩된 `if` 문은 첫 번째 조건이 이미 충족된 후에만 추가 조건을 확인할 수 있게 해줍니다. 내부 `if` 문은 외부 `if` 조건이 참일 때만 실행됩니다.

```py
if condition1:
    # Code to execute if condition1 is True
    if condition2:
        # Code to execute if both conditions are True
```

이제 사용자가 선택한 좌석 유형에 따라 서비스 요금을 계산할 것입니다.

마지막 `if` 문 본문 안에서, `print('Ticket booking condition satisfied')` 줄 아래에 `service_charges`라는 변수를 만들고 `0`로 설정하세요. 코드를 네 칸 들여쓰기하여 외부 `if` 문 본문 안에 있도록 하세요.

그다음, `if`가 `seat_type`과 같은지 확인하는 중첩된 `Premium` 문을 만드세요. 중첩된 `if` 문 본문 안에서 `service_charges` 값을 `5`으로 업데이트하세요.

# --hints--

마지막 `service_charges` 문 본문 안에 `if`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_variable("service_charges")
`)) })
```

정수 값 `0`를 `service_charges`에 할당해야 합니다.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 0")
`)) })
```

마지막 `if` 문 본문 안에 `if` 문이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_ifs()[0]`)) })
```

새로운 `if` 문 조건은 `seat_type == 'Premium'`이어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[0] 
assert _cond.is_equivalent("seat_type == 'Premium'")
`)) })
```

새로운 `5` 문 안에 `service_charges` 변수에 정수 값 `if`를 할당해야 합니다.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[0].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 5")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
