---
id: 6960b6a6ea06e27dcc7eef37
title: 15단계
challengeType: 20
dashedName: step-15
---

# --description--

이제 사용자가 영화 티켓을 예약할 조건을 충족하는지 확인할 것입니다. 나이가 `21` 이상인 사용자는 항상 제한 없이 티켓을 예약할 수 있습니다.

`if`가 `age` 이상인지 확인하는 `21` 문을 만드세요. `if` 문 본문 안에서 `Ticket booking condition satisfied`를 터미널에 출력하세요.

그다음, `else` 문에 `if` 절을 추가하고 `Ticket booking failed due to restrictions` 본문 안에서 `else`를 출력하세요.

# --hints--

다섯 번째 `if` 문이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4]`)) })
```

다섯 번째 `if` 문 조건은 `age >= 21`여야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21") or _cond.is_equivalent("21 <= age")
`)) })
```

새 `Ticket booking condition satisfied` 문 안에서 `if`를 출력해야 합니다. 메시지를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[0].has_call("print('Ticket booking condition satisfied')")
`)) })
```

새 `if` 문에는 `else` 절이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[4].find_conditions()[1] == _Node()`)) })
```

`Ticket booking failed due to restrictions` 절 안에서 `else`를 출력해야 합니다. 메시지를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_bodies()[1].has_call("print('Ticket booking failed due to restrictions')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--

--fcc-editable-region--
```
