---
id: 696099abb6acc539b27055a1
title: 5단계
challengeType: 20
dashedName: step-5
---

# --description--

이전 단계에서는 `if` 문으로 조건을 확인했습니다. `else` 절은 조건이 거짓일 때 처리하는 데 사용됩니다. `if...else` 문 구문은 다음과 같습니다:

```py
if condition:
    # Code to execute if condition is True
else:
    # Code to execute if condition is False
```

이제 `else` 문에 `if` 절을 추가하고 `User is not eligible for Evening shows` 본문 안에 `else` 를 출력하세요. 이렇게 하면 `if` 문 조건이 거짓일 때만 출력됩니다.

# --hints--

`else` 절이 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[1].find_conditions()[1].is_empty()`)) })
```

`User is not eligible for Evening shows` 절 안에 `else` 를 출력해야 합니다. 메시지는 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[1].find_bodies()[1].has_call("print('User is not eligible for Evening shows')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
--fcc-editable-region--

--fcc-editable-region--
```
