---
id: 6960a5165f7efe52a0c0aa39
title: 10단계
challengeType: 20
dashedName: step-10
---

# --description--

Python에서 `and` 연산자는 여러 조건이 모두 참인지 확인할 때 사용합니다. `if` 문에서 두 조건을 결합하는 방법은 다음과 같습니다.

```py
if condition1 and condition2:
    # Code to execute if all conditions are True
```

회원 할인은 `age`가 `21` 이상인 회원에게만 적용되어야 합니다.

`if is_member:` 줄의 조건을 `and` 연산자를 사용해 기존 조건과 `age`가 `21` 이상인지 확인하는 다른 조건을 결합하도록 업데이트하세요.

# --hints--

세 번째 `if` 문 조건은 `is_member and age >= 21`이어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[2].find_conditions()[0] 
assert _cond.is_equivalent("is_member and age >= 21") or _cond.is_equivalent("is_member and 21 <= age")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--
if is_member:
--fcc-editable-region--
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
