---
id: 6960ba6e6c95ed4ad896c6f2
title: 17단계
challengeType: 20
dashedName: step-17
---

# --description--

예약 규칙에는 한 가지 예외가 더 있습니다. `18`와 `21` 사이에 있는 사용자는 회원일 경우 저녁 공연을 예약할 수 있습니다.

`if` 문장에서 여러 논리 연산자가 사용될 때, `and`로 연결된 조건이 `or`로 연결된 조건보다 먼저 평가됩니다. Python에서는 소괄호 `()`를 사용해 조건을 그룹화하고 평가 순서를 제어합니다.

```py
if condition1 and (condition2 or condition3):
    # Code to execute if conditions evaluate to True
else:
    # Code to execute if conditions evaluate to False
```

`if age >= 21 or age >= 18 and show_time != 'Evening':` 줄의 조건을 업데이트하여 `or` 연산자를 사용해 `is_member`가 참인지 확인하는 또 다른 조건을 추가하세요. 위 예시처럼 `()` 소괄호를 사용해 `show_time != 'Evening'`과 `is_member` 조건을 함께 그룹화하세요.

# --hints--

마지막 `if` 문장 조건은 `age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)`이어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
assert _cond.is_equivalent("age >= 21 or age >= 18 and (show_time != 'Evening' or is_member)")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21 or age >= 18 and show_time != 'Evening':
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
