---
id: 6960b8632b9544d1f57295f4
title: 16단계
challengeType: 20
dashedName: step-16
---

# --description--

`18`와 `21` 사이의 사용자는 `show_time`가 `Evening`가 아닐 때만 티켓을 예약할 수 있습니다.

`if age >= 21:` 줄의 조건을 업데이트하세요. `and` 연산자를 사용해 `age`가 `18` 이상이고 `show_time`가 `Evening`가 아닌지 확인하는 표현식을 만드세요. 그런 다음 `or` 연산자를 사용해 그 표현식을 기존 조건과 결합하세요. 새 변수를 만들지 마세요.

# --hints--

마지막 `if` 문장 조건은 `age >= 21 or age >= 18 and show_time != 'Evening'`여야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_conditions()[0] 
_first = _cond.is_equivalent("age >= 21 or age >= 18 and show_time != 'Evening'")
_second = _cond.is_equivalent("age >= 21 or 18 <= age and show_time != 'Evening'")
assert _first or _second
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

--fcc-editable-region--
if age >= 21:
--fcc-editable-region--
    print('Ticket booking condition satisfied')
else:
    print('Ticket booking failed due to restrictions')
```
