---
id: 6960b308bf41cc39c9b4b438
title: 14단계
challengeType: 20
dashedName: step-14
---

# --description--

Python에서 `or` 연산자는 두 조건 중 적어도 하나가 참인지 확인하는 데 사용됩니다. `if` 문에서 이를 사용하는 방법은 다음과 같습니다.

```py
if condition1 or condition2:
    # Code to execute if any condition is True
```

공연이 저녁에 있을 경우 추가 요금도 적용되어야 합니다. `if is_weekend:` 줄의 조건을 `or` 연산자를 사용해 기존 조건과 두 번째 조건을 결합하도록 업데이트하세요. 두 번째 조건은 `show_time`가 문자열 `Evening`와 같은지 확인합니다.

# --hints--

마지막 `if` 문 조건은 `is_weekend or show_time == 'Evening'`이어야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[3].find_conditions()[0].is_equivalent("is_weekend or show_time == 'Evening'")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
--fcc-editable-region--
if is_weekend:
--fcc-editable-region--
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)
```
