---
id: 6960a848aacf859f7bb15447
title: 11단계
challengeType: 20
dashedName: step-11
---

# --description--

사용자가 회원이 아니므로 `is_member` 변수의 값을 `False`로 변경하세요.

그다음, 할인 적용을 위해 두 조건이 모두 충족되어야 하므로 `discount` 값이 이제 `0`로 유지되는 것을 확인할 수 있습니다.

# --hints--

`is_member` 변수의 값을 `False`로 업데이트하세요. 값에 따옴표를 붙이지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--
is_member = True
--fcc-editable-region--
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)
```
