---
id: 696099c1ac72a4e585de1d5b
title: 6단계
challengeType: 20
dashedName: step-6
---

# --description--

어떤 정보는 참 또는 거짓일 수밖에 없습니다. 이전 수업에서 배운 것처럼, 이는 불리언 값으로 나타낼 수 있습니다.

사용자가 회원인지 나타내기 위해 `is_member`라는 변수를 만들고 그 값을 `True`로 할당하세요.

`is_member` 변수 아래에 영화 상영이 주말인지 나타내는 `is_weekend`라는 변수를 만들고 그 값을 `False`로 할당하세요. 값은 따옴표로 감싸지 마세요.

# --hints--

`is_member`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_member")`)) })
```

불리언 값 `True`를 `is_member`에 할당해야 합니다. 값은 따옴표로 감싸지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_member").is_equivalent("is_member = True")
`)) })
```

`is_weekend`라는 변수를 선언해야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).has_variable("is_weekend")`)) })
```

불리언 값 `False`를 `is_weekend`에 할당해야 합니다. 값은 따옴표로 감싸지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_variable("is_weekend").is_equivalent("is_weekend = False")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

--fcc-editable-region--

--fcc-editable-region--
```
