---
id: 696cbb3985e23420cd45c58f
title: 20단계
challengeType: 20
dashedName: step-20
---

# --description--

`if...elif...else` 문장은 여러 조건을 순서대로 확인하는 데 사용됩니다.

```py
if condition1:
   # Code to execute if condition1 is True
elif condition2:
   # Code to execute if condition1 is False and condition2 is True
else:
   # Code to execute if all conditions are False
```

외부 `if` 문장의 본문 안에 있으면서, `elif`와 `if seat_type == 'Premium':` 줄 사이에 `else:` 절을 추가하고 `seat_type`가 `Gold`와 같은지 확인하세요. `elif` 절의 본문 안에서 `service_charges`의 값을 `3`로 업데이트하세요.

중첩된 `if...elif...else` 문장 아래에 `print()` 호출을 사용해 `Service charges:`와 업데이트된 `service_charges` 값을 보여주는 메시지를 출력하세요. 그런 다음 터미널에서 출력을 확인하세요.

# --hints--

`elif` 절이 있어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1]
assert not _cond.is_empty()
`)) })
```

`elif` 절 조건은 `seat_type == 'Gold'`이어야 합니다.

```js
({ test: () => (runPython(`
_cond = _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1] 
assert _cond.is_equivalent("seat_type == 'Gold'")
`)) })
```

`3` 본문 안에서 `service_charges` 변수에 정수 값 `elif`를 할당해야 합니다.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[1].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 3")
`)) })
```

`print('Service charges:', service_charges)` 문장 아래에 `if...elif...else`가 있어야 합니다.

```js
({ test: () => (runPython(`
if_stmt = """if seat_type == 'Premium':
    service_charges = 5
elif seat_type == 'Gold':
    service_charges = 3
else:
    service_charges = 1"""
print_call = "print('Service charges:', service_charges)"
assert _Node(_code).find_ifs()[4].find_bodies()[0].is_ordered(if_stmt, print_call)
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

    service_charges = 0
    if seat_type == 'Premium':
        service_charges = 5
--fcc-editable-region--
    
    else:
        service_charges = 1
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
