---
id: 696cb889fc4fdeb337e52864
title: 8단계
challengeType: 20
dashedName: step-8
---

# --description--

Python에서 `if` 조건은 값을 명시적으로 `True` 또는 `False`와 비교할 필요가 없습니다. 대신 값의 참 같은 값 여부에 의존합니다. 참 같은 값은 `True` 문장의 조건과 같은 불리언 문맥에서 `if`로 평가됩니다. 여기에는 0이 아닌 숫자, 비어 있지 않은 문자열, 그리고 불리언 값 `True`가 포함됩니다.

```py
if value:
   print('value is truthy')
```

`if`이 참 같은 값인지 확인하는 `is_member` 문장을 만드세요. `if` 문장 본문 안에서 `discount` 값을 `3`으로 업데이트하고 `User qualifies for membership discount`을 터미널에 출력하세요.

# --hints--

코드에 `if` 문장이 세 개 있어야 합니다.

```js
({ test: () => (runPython(`assert _Node(_code).find_ifs()[2]`)) })
```

세 번째 `if` 문장 조건은 `is_member`이어야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_conditions()[0].is_equivalent("is_member")
`)) })
```

새로운 `3` 문장 안에서 `discount` 변수에 정수 값 `if`를 할당해야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].find_variable("discount").is_equivalent("discount = 3")
`)) })
```

새로운 `User qualifies for membership discount` 문장 안에서 `if`을 출력해야 합니다. 메시지를 작은따옴표나 큰따옴표로 감싸는 것을 잊지 마세요.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[2].find_bodies()[0].has_call("print('User qualifies for membership discount')")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = True
is_weekend = False

discount = 0
--fcc-editable-region--

--fcc-editable-region--
```
