---
id: 6970f09b547e10402abf2847
title: 19단계
challengeType: 20
dashedName: step-19
---

# --description--

외부 `if` 문장 본문 안에 있으면서, 중첩된 `else` 문장에 `if seat_type == 'Premium':` 절을 추가하고 `service_charges` 본문 안에서 `1` 값을 `else`로 업데이트하세요. `else`가 외부 `if` 문장 본문 안에 있도록 네 칸 들여쓰기를 유지하세요.

# --hints--

새로운 `if` 문장에는 `else` 절이 있어야 합니다.

```js
({ test: () => (runPython(`
assert _Node(_code).find_ifs()[4].find_ifs()[0].find_conditions()[1].is_empty() 
`)) })
```

`1` 절 안에서 정수 값 `service_charges`을 `else` 변수에 할당해야 합니다.

```js
({ test: () => (runPython(`
_var = _Node(_code).find_ifs()[4].find_ifs()[0].find_bodies()[1].find_variable("service_charges") 
assert _var.is_equivalent("service_charges = 1")
`)) })
```

# --seed--

## --seed-contents--

```py
base_price = 15
age = 21
seat_type = 'Gold'
show_time = 'Evening'

if age > 17:
    print('User is eligible to book a ticket')

if age >= 21:
    print('User is eligible for Evening shows')
else:
    print('User is not eligible for Evening shows')

is_member = False
is_weekend = False

discount = 0
if is_member and age >= 21:
    discount = 3
    print('User qualifies for membership discount')
else:
    print('User does not qualify for membership discount')
print('Discount:', discount)

extra_charges = 0
if is_weekend or show_time == 'Evening':
    extra_charges = 2
    print('Extra charges will be applied')
else:
    print('No extra charges will be applied')
print('Extra charges:', extra_charges)

if age >= 21 or age >= 18 and (show_time != 'Evening' or is_member):
    print('Ticket booking condition satisfied')

    service_charges = 0
    if seat_type == 'Premium':
        service_charges = 5
--fcc-editable-region--
    
--fcc-editable-region--
else:
    print('Ticket booking failed due to restrictions')
```
