---
id: 66da326c02141df538f29ba5
title: 创建一个结账页面
challengeType: 25
dashedName: build-a-checkout-page
demoType: onClick
---

# --description--

**目标：**实现以下用户需求并通过所有测试以完成实验。

**用户需求：**

1. 你应该有一个 `h1` 元素，文本为 `Checkout` 。
1. `h1` 元素之后应紧跟两个 `section` 元素。
1. 在第一部分中，应包含一个 `h2` 元素，文本为 `Your Cart`。
1. 应在第一部分提供项目图片，并配以适当的备用文字。 如果你愿意，可以使用这张图片：`https://cdn.freecodecamp.org/curriculum/labs/cube.jpg`
1. 在第二部分中，应包含一个 `h2` 元素，文本为 `Payment Information`。
1. 在第二部分中，应有一个 `form` 元素。
1. 在表单中，应包含一个输入元素，其 `id` 和 `name` 属性均为 `card-name`，并且有与之关联的 `label`。
1. 在表单中，应包含一个输入元素，其 `id` 和 `name` 属性均为 `card-number`，并且有与之关联的 `label`。
1. 你至少应有两个输入元素是必填项，并且应将 `aria-required` 属性设置为 `true`。

# --hints--

你应该有一个 `h1` 元素，文本为 `Checkout` 。

```js
assert.equal(document.querySelector('h1')?.innerText, 'Checkout');
```

你的页面上应该只有一个 `h1` 元素。

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

在 `h1` 元素之后，你应至少包含两个 `section` 元素。

```js
const h1 = document.querySelector('h1');
const firstSibling = h1?.nextElementSibling;
const secondSibling = firstSibling?.nextElementSibling;
assert.equal(firstSibling?.tagName, 'SECTION');
assert.equal(secondSibling?.tagName, 'SECTION');
```

在第一部分中，应包含一个 `h2` 元素，文本为 `Your Cart`。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.equal(firstSection?.querySelector('h2')?.innerText, 'Your Cart');
```

在第一部分中，应该至少有一张带有替代文本的图片。

```js
const firstSection = document.querySelectorAll('section')?.[0];
assert.isAtLeast(firstSection?.querySelector('img')?.alt?.length, 1);
```

在第二部分中，应包含一个 `h2` 元素，文本为 `Payment Information`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.equal(secondSection?.querySelector('h2')?.innerText, 'Payment Information');
```

在第二个 `section` 中，应包含一个 `form` 元素。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form'));
```

在表单中，应包含一个 `input` 元素，其 `id` 和 `name` 为 `card-name`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-name'));
assert.exists(secondSection?.querySelector('form input[name="card-name"]'));
```

在表单中，应包含一个 `input` 元素，其 `id` 和 `name` 为 `card-number`。

```js
const secondSection = document.querySelectorAll('section')?.[1];
assert.exists(secondSection?.querySelector('form input#card-number'));
assert.exists(secondSection?.querySelector('form input[name="card-number"]'));
```

所有非 `submit` `type` 的 `input` 元素，都应有一个与之关联的 `label` 元素。

```js
const inputs = document.querySelectorAll('input:not([type="submit"])');
assert.isAtLeast(inputs.length, 1);
inputs.forEach(input => assert.exists(document.querySelector(`label[for="${input.id}"]`)));
```

你至少应有两个 `input` 元素是必填项，并且应将 `aria-required` 属性设置为 `true`。

```js
const inputs = document.querySelectorAll('input');
const requiredInputs = Array.from(inputs).filter(input => input.required && input.getAttribute('aria-required') === 'true');
assert.isAtLeast(requiredInputs?.length, 2)
```


# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8" />
    <title>Checkout Page</title>
</head>

<body>

</body>

</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Checkout Page</title>
</head>
<body>
    <div class="container">
        <h1>Checkout</h1>

        <section>
            <h2>Your Cart</h2>
            <img src="https://cdn.freecodecamp.org/curriculum/labs/spoon.jpg" alt="Left-handed spoon">
            <p>Left-handed spoon</p>
            <p>$29.99</p>
        </section>

        <section>
            <h2>Payment Information</h2>
            <form action="/submit-payment" method="POST">
                <div>
                    <label for="card-name">Cardholder Name</label>
                    <input type="text" id="card-name" name="card-name" required aria-required="true">
                </div>
                <div>
                    <label for="card-number">Card Number</label>
                    <input type="text" id="card-number" name="card-number" required aria-required="true">
                </div>
                <div>
                    <label for="expiry-date">Expiry Date</label>
                    <input type="text" id="expiry-date" name="expiry-date" placeholder="MM/YY" required aria-required="true">
                </div>
                <div>
                    <label for="cvv">CVV</label>
                    <input type="text" id="cvv" name="cvv" required aria-required="true" aria-label="Card Verification Value">
                </div>
                <input type="submit" value="Place Order">
            </form>
        </section>
    </div>
</body>
</html>
```
