---
id: 66c06d618d075c7f7f1b890a
title: Construye una máquina de adivinación
challengeType: 26
dashedName: build-a-fortune-teller
---

# --description--

En este laboratorio, definirás cinco fortunas y seleccionarás aleatoriamente una de ellas para mostrarla al usuario.

**Objective:** Fulfill the user stories below and get all the tests to pass to complete the lab.

**Historia de usuario:**

1. Debes inicializar las cinco variables `fortune1`, `fortune2`, `fortune3`, `fortune4` y `fortune5` con un valor de cadena de tu elección. Puedes usar las siguientes opciones si lo deseas:

   - `"Your cat will look very cuddly today."`
   - `"The weather will be nice tomorrow."`
   - `"Be cautious of your new neighbors."`
   - `"You will find a new hobby soon."`
   - `"It would be wise to avoid the color red today."`

2. Debes seleccionar un número aleatorio entre 1 y 5, inclusive, y asignarlo a la variable `randomNumber`.

3. Debes crear una variable `selectedFortune` y asignar la fortuna adecuada según estas reglas:

   - Si `randomNumber` es 1, asigna el valor de `fortune1` a `selectedFortune`.
   - Si `randomNumber` es 2, asigna el valor de `fortune2` a `selectedFortune`.
   - Si `randomNumber` es 3, asigna el valor de `fortune3` a `selectedFortune`.
   - Si `randomNumber` es 4, asigna el valor de `fortune4` a `selectedFortune`.
   - Si `randomNumber` es 5, asigna el valor de `fortune5` a `selectedFortune`.

4. Debes registrar en la consola el `selectedFortune`.

# --hints--

Debes inicializar `fortune1` con un valor de cadena.

```js
assert.isNotNull(fortune1);
assert.isString(fortune1);
```

Debes inicializar `fortune2` con un valor de cadena.

```js
assert.isNotNull(fortune2);
assert.isString(fortune2);
```

Debes inicializar `fortune3` con un valor de cadena.

```js
assert.isNotNull(fortune3);
assert.isString(fortune3);
```

Debes inicializar `fortune4` con un valor de cadena.

```js
assert.isNotNull(fortune4);
assert.isString(fortune4);
```

Debes inicializar `fortune5` con un valor de cadena.

```js
assert.isNotNull(fortune5);
assert.isString(fortune5);
```

Debes usar el método `Math.random()` para generar un número aleatorio.

```js
assert.match(__helpers.removeJSComments(code), /Math\.random\(\)/);
```

Debes generar un número aleatorio entre 1 y 5, inclusive, y asignarlo a la variable `randomNumber`.

```js
assert.isNotNull(randomNumber);
assert.include([1,2,3,4,5], randomNumber);
```

Debes tener una variable `selectedFortune` que se asigne un valor basado en el valor de `randomNumber`.

```js
assert.isNotNull(selectedFortune);
```

El `randomNumber` debe corresponder a su fortuna. Por ejemplo, si `randomNumber` es 1, el `selectedFortune` debe ser igual a `fortune1`, y así sucesivamente.

```js
const condition1 = randomNumber === 1 && selectedFortune === fortune1;
const condition2 = randomNumber === 2 && selectedFortune === fortune2;
const condition3 = randomNumber === 3 && selectedFortune === fortune3;
const condition4 = randomNumber === 4 && selectedFortune === fortune4;
const condition5 = randomNumber === 5 && selectedFortune === fortune5;

assert.isTrue(condition1 || condition2 || condition3 || condition4 || condition5);
```

Debes enviar el `selectedFortune` a la consola.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*selectedFortune\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const fortune1 = "Your cat will look very cuddly today.";
const fortune2 = "The weather will be nice tomorrow.";
const fortune3 = "Be cautious of your new neighbors.";
const fortune4 = "You will find a new hobby soon.";
const fortune5 = "It would be wise to avoid the color red today.";

let randomNumber = Math.floor(Math.random() * 5) + 1;

let selectedFortune;

if (randomNumber === 1) {
  selectedFortune = fortune1;
} else if (randomNumber === 2) {
  selectedFortune = fortune2;
} else if (randomNumber === 3) {
  selectedFortune = fortune3;
} else if (randomNumber === 4) {
  selectedFortune = fortune4;
} else {
  selectedFortune = fortune5;
}

console.log(selectedFortune);
```
