---
id: 655cd899f8de09431eabb40c
title: الخطوة 1
challengeType: 20
dashedName: step-1
---

# --description--

فرز الدمج هو خوارزمية فرز تستخدم مبدأ القسمة والتغلب لفرز مجموعات البيانات. أي أنها "تقسم" المجموعة إلى أجزاء فرعية أصغر، و"تتغلب" على الأجزاء الفرعية بفرزها بشكل مستقل، ثم تدمج الأجزاء الفرعية المرتبة.

ابدأ بإنشاء دالة تسمى `merge_sort` تأخذ المعلمة `array`. ستتولى هذه الدالة مهمة فرز قائمة من الأرقام. استخدم الكلمة المفتاحية `pass` داخل جسم الدالة.

# --hints--

يجب إعلان دالة باسم `merge_sort` بمعلمة واحدة فقط: `array`. لا تنسَ استخدام الكلمة المفتاحية `pass` داخل جسم الدالة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_args("array")`))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
