---
id: 6564aee9c077774ea49c3faf
title: الخطوة 2
challengeType: 20
dashedName: step-2
---

# --description--

تُنفذ خوارزمية دمج الترتيب ثلاث عمليات:

- قسّم متتالية غير مرتبة من العناصر إلى أجزاء فرعية
- رتب العناصر في الأجزاء الفرعية
- دمج الأجزاء الفرعية المرتبة

تحدث العمليات السابقة بشكل متكرر حتى تدمج الأجزاء الفرعية في المتتالية المرتبة الكاملة. ابدأ بتقسيم المتتالية.

أولًا، استبدل الكلمة المفتاحية `pass` بمتغير `middle_point`. استخدم عامل القسمة الصحيحة (`//`) لقسمة طول قائمة `array` على 2 وعيّن النتيجة في متغير `middle_point` الجديد. تذكر أن تضع الشفرة بمستوى المسافة البادئة الصحيح.

# --hints--

يجب ألا يحتوي جسم الدالة على `pass`.

```js
({
  test: () => assert.isFalse(runPython(`_Node(_code).find_function("merge_sort").has_pass()`))
})
```

يجب إعلان متغير `middle_point` وتعيين `len(array) // 2` له.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_stmt("middle_point = len(array) // 2")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
