---
id: 656627b47bd2d2a4afbc945d
title: الخطوة 3
challengeType: 20
dashedName: step-3
---

# --description--

في الخطوة السابقة، حصلت على نقطة المنتصف. يمكنك استخدامها لتقسيم `array` إلى قسمين وتعيين كل جزء إلى متغير جديد.

استخدم صيغة القطع لاستخراج النصف الأيسر من `array` وتعيينه إلى متغير باسم `left_part`.

# --hints--

يجب أن يكون لديك متغير باسم `left_part` في دالة `merge_sort` الخاصة بك.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_part")`)) 
})
```

يجب أن تستخدم صيغة القطع لاستخراج النصف الأيسر من المصفوفة وتعيينه إلى المتغير `left_part`.

```js
({
  test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_variable("left_part")
    cond.is_equivalent("left_part = array[:middle_point]") or cond.is_equivalent("left_part = array[0:middle_point]")
  `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
--fcc-editable-region--
    
--fcc-editable-region--
```
