---
id: 656639771fed09b5c0e8fe71
title: الخطوة 5
challengeType: 20
dashedName: step-5
---

# --description--

بعد أن قسمت قائمة `array` إلى قائمتين منفصلتين، استمر في تقسيم كل قائمة حتى يصبح كل عنصر بمفرده في قائمته الخاصة. القائمة التي تحتوي على رقم واحد دائمًا ما تكون مرتبة.

للقيام بذلك، استدعِ دالة `merge_sort` بشكل متكرر داخل دالتك ومرر `left_part` كمعلمة في الاستدعاء.

# --hints--

يجب أن تستدعي `merge_sort` مع المعلمة `left_part` في أسفل جسم الدالة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("right_part = array[middle_point:]", "merge_sort(left_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

--fcc-editable-region--
```
