---
id: 656680b0fc79f2c38a34d90e
title: الخطوة 7
challengeType: 20
dashedName: step-7
---

# --description--

حان الوقت الآن لفرز ودمج القوائم (`left_part` و `right_part`) في القائمة الأصلية.

يمكنك فعل ذلك بمقارنة العناصر في كلتا القائمتين، ودمج العنصر الأصغر في القائمة الرئيسية. ستقوم بهذه المقارنة لكل الفهارس في `left_part` و `right_part`.

أنشئ ثلاث متغيرات: `left_array_index` و `right_array_index` و `sorted_index` وضع قيمها على `0`. ستساعدك هذه المتغيرات على تتبع كل فهرس أثناء عملية الفرز.

# --hints--

يجب أن يكون لديك متغير باسم `left_array_index` داخل دالة `merge_sort` الخاصة بك

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("left_array_index")`)) 

})
```

يجب أن تكون قيمة `left_array_index` هي `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("left_array_index").is_equivalent("left_array_index = 0")`))  

})
```

يجب أن يكون لديك متغير باسم `right_array_index` داخل دالة `merge_sort` الخاصة بك

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_array_index")`))
})
```

يجب أن تكون قيمة `right_array_index` هي `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_array_index").is_equivalent("right_array_index = 0")`))
})

```

يجب أن يكون لديك متغير باسم `sorted_index` داخل دالة `merge_sort` الخاصة بك

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("sorted_index")`))
})
```

يجب أن تكون قيمة `sorted_index` هي `0`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("sorted_index").is_equivalent("sorted_index = 0")`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
--fcc-editable-region--
    
--fcc-editable-region--
```
