---
id: 656702f8b4cbd8cbf0a433c6
title: الخطوة 8
challengeType: 20
dashedName: step-8
---

# --description--

بعد ذلك، عليك إنشاء حلقة تستمر طالما توجد عناصر متبقية في كل من `left_part` و `right_part`.

لهذا، أنشئ حلقة `while` بشرطين: الأول يتحقق مما إذا كان `left_array_index` أقل من طول `left_part`، والآخر يتحقق مما إذا كان `right_array_index` أقل من طول `right_part`.

# --hints--

يجب أن تتحقق حلقة `while` من كون `left_array_index` أقل من طول `left_part`، ومن كون `right_array_index` أقل من طول `right_part`.

```js
({ test: () => assert(runPython(`
    cond= _Node(_code).find_function("merge_sort").find_whiles()[0].find_conditions()[0]

    cond.is_equivalent("left_array_index < len(left_part) and right_array_index < len(right_part)") or cond.is_equivalent("right_array_index < len(right_part) and left_array_index < len(left_part)")
    `))
}) 
```

استخدم الكلمة المفتاحية `pass` داخل جسم الحلقة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[0].find_body().has_pass()`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
