---
id: 656709e50ed928da35260276
title: الخطوة 12
challengeType: 20
dashedName: step-12
---

# --description--

في خطوة سابقة، عيّنت العنصر في `left_part` إلى قائمة `array` لأنه كان أصغر. لكن هذا لن يكون الحال دائمًا. في بعض حالات المقارنة، قد يكون العنصر على اليمين أصغر.

أنشئ بند `else` للتنفيذ عندما لا يكون العنصر في `left_part` أقل من العنصر في `right_part`.

داخل كتلة `else`، عيّن `right_part[right_array_index]` إلى `array[sorted_index]`.

# --hints--

يجب أن يحتوي جسم بند `else` على `array[sorted_index] = right_part[right_array_index]`.

```js
({
  test: () => assert(runPython(`
  _Node(_code).find_function("merge_sort").find_whiles()[0].find_ifs()[0].find_bodies()[1].find_body()[0].is_equivalent("array[sorted_index] = right_part[right_array_index]")
`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
--fcc-editable-region--
        
--fcc-editable-region--
```
