---
id: 656710d1e0ec62253426db24
title: الخطوة 15
challengeType: 20
dashedName: step-15
---

# --description--

حلقة `while` التي أنشأتها تقارن عنصرًا واحدًا من `left_part` مع عنصر آخر في `right_part`، ثم تضيف العنصر الأصغر إلى القائمة الرئيسية `array`.

ستستمر هذه العملية حتى لا يبقى عناصر للمقارنة. لكن قد تبقى عناصر في `left_part` بينما لا يوجد في `right_part`، والعكس صحيح.

أنشئ حلقة `while` أخرى تعمل عندما يكون `left_array_index` أقل من `len(left_part)`. في الخطوات التالية، ستستخدمها لنسخ العناصر المتبقية في `left_part` إلى القائمة `array`.

# --hints--

يجب أن تحتوي حلقة `while` على هذا الشرط: `left_array_index < len(left_part)`. لا تنسَ استخدام الكلمة المفتاحية `pass` في جسم الحلقة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_conditions()[0].is_equivalent("left_array_index < len(left_part)")`))
})
```

يجب أن تحتوي حلقة `while` على الكلمة المفتاحية `pass` في جسمها

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[1].find_bodies()[0].has_pass()`))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
