---
id: 6569b4e0bd29d17d4b53b16c
title: الخطوة 19
challengeType: 20
dashedName: step-19
---

# --description--

الآن، ستكرر نفس منطق حلقة `while` لـ `right_part`.

أنشئ حلقة `while` تعمل طالما أن `right_array_index` أقل من `len(right_part)`. داخل حلقة الـ `while`، قم بتعيين `right_part[right_array_index]` إلى `array[sorted_index]`. في النهاية، زد كل من `right_array_index` و`sorted_index` بمقدار `1`.

# --hints--

يجب أن تتحقق حلقة `while` من أن `right_array_index` أقل من `len(right_part)`.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_conditions()[0].is_equivalent("right_array_index < len(right_part)")`)) 
})
```

يجب أن تعيّن `right_part[right_array_index]` إلى `array[sorted_index]` داخل جسم الحلقة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_whiles()[2].find_body().has_stmt("array[sorted_index] = right_part[right_array_index]")`))
})
```

يجب زيادة قيمة `right_array_index` بمقدار `1` داخل حلقة الـ `while`. لاحظ أن الزيادة يجب أن تتم بعد تعيين عنصر الجزء الأيمن إلى المصفوفة المرتبة.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = right_array_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "right_array_index = 1 + right_array_index")
    `))
}) 
```

يجب زيادة القيمة الحالية لـ `sorted_index` بمقدار `1` داخل حلقة الـ `while`. لاحظ أن الزيادة يجب أن تتم بعد تعيين عنصر الجزء الأيمن إلى المصفوفة المرتبة.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[2].find_body()
    
    cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = right_part[right_array_index]", "sorted_index = 1 + sorted_index")
    `))
}) 
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1

    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--
    
--fcc-editable-region--
```
