---
id: 6569b20f829b7e69d43c232a
title: الخطوة 18
challengeType: 20
dashedName: step-18
---

# --description--

آخر شيء يجب فعله في حلقة `while` هو الانتقال إلى الفهرس التالي في المصفوفة المرتبة.

أضف `1` إلى قيمة `sorted_index`.

# --hints--

يجب زيادة القيمة الحالية لـ `sorted_index` بمقدار واحد داخل حلقة `while`. لاحظ أن الزيادة يجب أن تتم بعد تعيين عنصر الجزء الأيسر إلى المصفوفة المرتبة.

```js
({ test: () => assert(runPython(`
    cond = _Node(_code).find_function("merge_sort").find_whiles()[1].find_body()
    
    cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index += 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = sorted_index + 1") or cond.is_ordered("array[sorted_index] = left_part[left_array_index]", "sorted_index = 1 + sorted_index") 
    `))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]

    merge_sort(left_part)
    merge_sort(right_part)

    left_array_index = 0
    right_array_index = 0
    sorted_index = 0

    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
        
--fcc-editable-region--
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1

--fcc-editable-region--
```
