---
id: 6577254891048c8f2c19e961
title: الخطوة 6
challengeType: 20
dashedName: step-6
---

# --description--

في أسفل جسم الدالة، استدعِ الدالة `merge_sort()` مرة أخرى. هذه المرة، مرر `right_part` كمعلمة لاستدعاء الدالة.

# --hints--

يجب أن تستدعي `merge_sort()` مع المعلمة `right_part` في أسفل جسم الدالة.

```js
({
    test: () => assert(runPython(`_Node(_code).find_function("merge_sort").is_ordered("merge_sort(left_part)", "merge_sort(right_part)")`
))
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
--fcc-editable-region--
    
--fcc-editable-region--
```
