---
id: 657f59751d5a5c9b069ae0f3
title: الخطوة 21
challengeType: 20
dashedName: step-21
---

# --description--

كما تعلمت في درس سابق، يمكنك استخدام المتغير `__name__` لتحديد ما إذا كان سكربت Python يُنفذ كبرنامج رئيسي أو إذا كان يُستورد كوحدة (شفرة مكتوبة في ملف Python آخر).

إذا كانت قيمة `__name__` مضبوطة على `'__main__'`، فهذا يعني أن السكربت الحالي هو البرنامج الرئيسي، وليس وحدة.

في هذا المشروع، ستستخدم السكربت الحالي كبرنامج رئيسي.

أنشئ جملة `if` تتحقق مما إذا كانت قيمة `__name__` هي `'__main__'`.
داخل جملة `if`، أنشئ قائمة باسم `numbers`، وقم بتعيين هذه القيم لها: `[4, 10, 6, 14, 2, 1, 8, 5]`.

# --hints--

يجب أن تتحقق جملة `if` الخاصة بك مما إذا كانت قيمة `__name__` هي `'__main__'`

```js
({
    test: () => assert(runPython(`
      _Node(_code).find_ifs()[0].find_conditions()[0].is_equivalent("__name__ == '__main__'")
    `))
})
```

يجب أن تحتوي قائمة `numbers` على هذه القيم: `[4, 10, 6, 14, 2, 1, 8, 5]`

```js
({
  test: () => assert(runPython(`_Node(_code).find_ifs()[0].find_body().find_variable("numbers").is_equivalent("numbers = [4, 10, 6, 14, 2, 1, 8, 5]")`)) 
})
```

# --seed--

## --seed-contents--

```py
def merge_sort(array):
    if len(array) <= 1:
        return
    
    middle_point = len(array) // 2
    left_part = array[:middle_point]
    right_part = array[middle_point:]
    
    merge_sort(left_part)
    merge_sort(right_part)
    
    left_array_index = 0
    right_array_index = 0
    sorted_index = 0
    
    while left_array_index < len(left_part) and right_array_index < len(right_part):
        if left_part[left_array_index] < right_part[right_array_index]:
            array[sorted_index] = left_part[left_array_index]
            left_array_index += 1
        else:
            array[sorted_index] = right_part[right_array_index]
            right_array_index += 1
        sorted_index += 1
    
    while left_array_index < len(left_part):
        array[sorted_index] = left_part[left_array_index]
        left_array_index += 1
        sorted_index += 1
    
    while right_array_index < len(right_part):
        array[sorted_index] = right_part[right_array_index]
        right_array_index += 1
        sorted_index += 1
    
--fcc-editable-region--

--fcc-editable-region--
```
