---
id: 657f3a661730891aa64f3568
title: الخطوة 4
challengeType: 20
dashedName: step-4
---

# --description--

استخدم صيغة القطع لاستخراج النصف الأيمن من `array` وتعيينه إلى متغير باسم `right_part`.

# --hints--

يجب أن يكون لديك متغير باسم `right_part`

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").has_variable("right_part")`))
})
```

يجب أن تستخدم صيغة القطع لاستخراج النصف الأيمن من المصفوفة وتعيينه إلى المتغير `right_part`.

```js
({
  test: () => assert(runPython(`_Node(_code).find_function("merge_sort").find_variable("right_part").is_equivalent("right_part = array[middle_point:]")`)) 
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
def merge_sort(array):
    middle_point = len(array) // 2
    left_part = array[:middle_point]
--fcc-editable-region--
```
