---
id: 613297a923965e0703b64796
title: Крок 2
challengeType: 0
dashedName: step-2
---

# --description--

Ви, можливо, вже знайомі з елементом `meta`; він використовується для вказівки інформації про сторінку, такої як заголовок, опис, ключові слова та автор.

Додайте до своєї сторінки елемент `meta` з відповідним значенням `charset`.

Атрибут `charset` визначає кодування символів сторінки, і сьогодні `UTF-8` — це єдине кодування, яке підтримують більшість браузерів.

# --hints--

Ви повинні створити елемент `meta` всередині елемента `head`.

```js
assert.exists(document.querySelector('head > meta'));
```

Ви повинні надати тегу `meta` атрибут `charset` зі значенням `UTF-8`.

```js
assert.equal(document.querySelector('head > meta')?.getAttribute('charset')?.toLowerCase(), 'utf-8');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
--fcc-editable-region--
  <head>
    <link rel="stylesheet" href="styles.css" />
  </head>
--fcc-editable-region--
  <body>

  </body>
</html>

```

```css
body {
  background: #f5f6f7;
  color: #1b1b32;
  font-family: Helvetica;
  margin: 0;
}
```
