---
id: 68e51d8f163d3f8b70f9c4e5
title: Passo 1
challengeType: 1
dashedName: step-1
---

# --description--

Neste workshop, você praticará o uso dos métodos `includes()` e `slice()`.

Para começar, crie uma variável chamada `fccSentence` e atribua a ela a string `"freeCodeCamp is a great place to learn web development."`.

# --hints--

Você deve declarar uma variável chamada `fccSentence`.

```js
assert.exists(fccSentence);
```

Sua variável `fccSentence` deve ser uma string.

```js
assert.isString(fccSentence);
```

Sua variável `fccSentence` deve ser atribuída à string `"freeCodeCamp is a great place to learn web development."`.

```js
assert.strictEqual(fccSentence, "freeCodeCamp is a great place to learn web development.");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
