---
id: 68e6172975008fea742a63d1
title: Passo 11
challengeType: 1
dashedName: step-11
---

# --description--

O método `slice()` pode usar índices negativos para contar a partir do final da string. `-1` se refere ao último caractere.

Aqui está um exemplo:

```js
const text = "JavaScript";
const lastThree = text.slice(-3);
console.log(lastThree); // "ipt"
```

Crie uma variável chamada `endPunctuation`. Depois, atribua a ela o resultado de usar `message.slice(-1)` para extrair o último caractere.

Então use um template literal ou concatenação de strings para registrar a mensagem `The ending punctuation mark is a "<endPunctuation>"` no console. Substitua `<endPunctuation>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve declarar uma variável chamada `endPunctuation`.

```js
assert.exists(endPunctuation);
```

Sua variável `endPunctuation` deve ser uma string.

```js
assert.isString(endPunctuation);
```

Você deve usar `message.slice(-1)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*-1\s*\)/);
```

Você deve atribuir o resultado de `message.slice(-1)` à sua variável `endPunctuation`.

```js
assert.match(__helpers.removeJSComments(code), /endPunctuation\s*=\s*message\.slice\(\s*-1\s*\)/);
```

Sua variável `endPunctuation` deve ser igual a `"!"`.

```js
assert.strictEqual(endPunctuation, "!");
```

Você deve registrar a mensagem `The ending punctuation mark is a "<endPunctuation>"`, onde `<endPunctuation>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['The ending punctuation mark is a "!"']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

const greetingWord = message.slice(0, 7);
console.log(`The first word is "${greetingWord}".`);

--fcc-editable-region--

--fcc-editable-region--
```
