---
id: 68e6172975008fea742a63d0
title: Passo 10
challengeType: 1
dashedName: step-10
---

# --description--

Você pode usar `slice()` para extrair do início de uma string.

Crie uma variável chamada `greetingWord`. Então atribua a ela o resultado de usar `message.slice(0, 7)` para extrair a primeira palavra `"Welcome"`.

Então use um template literal ou concatenação de strings para registrar a mensagem `The first word is "<greetingWord>".` no console. Substitua `<greetingWord>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve declarar uma variável chamada `greetingWord`.

```js
assert.exists(greetingWord);
```

Sua variável `greetingWord` deve ser uma string.

```js
assert.isString(greetingWord);
```

Você deve usar `message.slice(0, 7)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Você deve atribuir o resultado de `message.slice(0, 7)` à sua variável `greetingWord`.

```js
assert.match(__helpers.removeJSComments(code), /greetingWord\s*=\s*message\.slice\(\s*0\s*,\s*7\s*\)/);
```

Sua variável `greetingWord` deve ser igual a `"Welcome"`.

```js
assert.strictEqual(greetingWord, "Welcome");
```

Você deve registrar a mensagem `The first word is "<greetingWord>".`, onde `<greetingWord>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['The first word is "Welcome".']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
console.log(`The word "${platform}" was sliced from the message.`);

--fcc-editable-region--

--fcc-editable-region--
```
