---
id: 68e6172975008fea742a63cf
title: Passo 9
challengeType: 1
dashedName: step-9
---

# --description--

Agora é hora de registrar o resultado do uso do método `slice()`.

Use uma template literal ou concatenação de strings para registrar a mensagem `The word "<platform>" was sliced from the message.` no console. Substitua `<platform>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve registrar a mensagem `The word "<platform>" was sliced from the message.`, onde `<platform>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['The word "freeCodeCamp" was sliced from the message.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

const platform = message.slice(11, 23);
--fcc-editable-region--

--fcc-editable-region--
```
