---
id: 68e6172975008fea742a63cc
title: Passo 6
challengeType: 1
dashedName: step-6
---

# --description--

O método `includes()` diferencia maiúsculas de minúsculas, ou seja, ele distingue entre letras maiúsculas e minúsculas.

Crie uma variável chamada `hasLowercaseFCC`. Depois, atribua a ela o resultado do uso do método `includes()` para verificar se `fccSentence` inclui `"freecodecamp"` (tudo em minúsculas).

Então use um template literal ou concatenação de strings para registrar a mensagem `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.` no console. Substitua `<hasLowercaseFCC>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve declarar uma variável chamada `hasLowercaseFCC`.

```js
assert.exists(hasLowercaseFCC);
```

Você deve usar `fccSentence.includes("freecodecamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

Você deve atribuir o resultado de `fccSentence.includes("freecodecamp")` à sua variável `hasLowercaseFCC`.

```js
assert.match(__helpers.removeJSComments(code), /hasLowercaseFCC\s*=\s*fccSentence\.includes\(\s*(['"`])freecodecamp\1\s*\)/);
```

O resultado de `fccSentence.includes("freecodecamp")` deve ser `false`. 

```js
assert.isFalse(hasLowercaseFCC);
```

Você deve registrar a mensagem que inclui `fccSentence.includes("freecodecamp") returns <hasLowercaseFCC> because includes is case-sensitive.`, onde `<hasLowercaseFCC>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freecodecamp") returns false because includes is case-sensitive.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
