---
id: 68e6172975008fea742a63cb
title: Passo 5
challengeType: 1
dashedName: step-5
---

# --description--

Agora é hora de ver o que acontece quando a substring NÃO é encontrada.

Crie uma variável chamada `hasJavaScript`. Então atribua a ela o resultado do uso do método `includes()` para verificar se `fccSentence` inclui `"JavaScript"`.

Então use uma template literal ou concatenação de strings para registrar a mensagem `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.` no console. Substitua `<hasJavaScript>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve declarar uma variável chamada `hasJavaScript`.

```js
assert.exists(hasJavaScript);
```

Você deve usar `fccSentence.includes("JavaScript")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

Você deve atribuir o resultado de `fccSentence.includes("JavaScript")` à sua variável `hasJavaScript`.

```js
assert.match(__helpers.removeJSComments(code), /hasJavaScript\s*=\s*fccSentence\.includes\(\s*(['"`])JavaScript\1\s*\)/);
```

O resultado de `fccSentence.includes("JavaScript")` deve ser `false`. 

```js
assert.isFalse(hasJavaScript);
```

Você deve registrar a mensagem `fccSentence.includes("JavaScript") returns <hasJavaScript> because the word "JavaScript" is not in the sentence.`, onde `<hasJavaScript>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("JavaScript") returns false because the word "JavaScript" is not in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

--fcc-editable-region--

--fcc-editable-region--
```
