---
id: 68e6172975008fea742a63cd
title: Passo 7
challengeType: 1
dashedName: step-7
---

# --description--

Agora você vai explorar o método `slice()`, que extrai uma parte de uma string.

Crie uma variável chamada `message` e atribua a ela a string `"Welcome to freeCodeCamp!"`.

Então use `console.log()` para registrar `"Here are some examples of the slice() method:"` no console.

# --hints--

Você deve declarar uma variável chamada `message`.

```js
assert.exists(message);
```

Sua variável `message` deve ser uma string.

```js
assert.isString(message);
```

Sua variável `message` deve ser `"Welcome to freeCodeCamp!"`.

```js
assert.strictEqual(message, "Welcome to freeCodeCamp!");
```

Você deve registrar `"Here are some examples of the slice() method:"` no console.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\s*\(\s*('|"|`)Here are some examples of the slice\(\) method:\1\s*\)/);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

--fcc-editable-region--

--fcc-editable-region--
```
