---
id: 68e6172975008fea742a63ce
title: Passo 8
challengeType: 1
dashedName: step-8
---

# --description--

O método `slice()` extrai uma seção de uma string. Ele recebe dois argumentos: o índice inicial (incluído) e o índice final (não incluído).

Aqui está um exemplo:

```js
const greeting = "Hello World";
const firstWord = greeting.slice(0, 5);
console.log(firstWord); // "Hello"
```

Os índices de string começam em 0. Na string `message`, o caractere no índice 11 é `f` (o início de "freeCodeCamp").

Crie uma variável chamada `platform`. Depois atribua a ela o resultado de usar `message.slice(11, 23)` para extrair `"freeCodeCamp"`.

# --hints--

Você deve declarar uma variável chamada `platform`.

```js
assert.exists(platform);
```

Sua variável `platform` deve ser uma string.

```js
assert.isString(platform);
```

Você deve usar `message.slice(11, 23)`.

```js
assert.match(__helpers.removeJSComments(code), /message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Você deve atribuir o resultado de `message.slice(11, 23)` à sua variável `platform`.

```js
assert.match(__helpers.removeJSComments(code), /platform\s*=\s*message\.slice\(\s*11\s*,\s*23\s*\)/);
```

Sua variável `platform` deve ser igual a `"freeCodeCamp"`.

```js
assert.strictEqual(platform, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
console.log(`fccSentence.includes("freeCodeCamp") returns ${hasFreeCodeCamp} because the word "freeCodeCamp" is in the sentence.`);

const hasJavaScript = fccSentence.includes("JavaScript");
console.log(`fccSentence.includes("JavaScript") returns ${hasJavaScript} because the word "JavaScript" is not in the sentence.`);

const hasLowercaseFCC = fccSentence.includes("freecodecamp");
console.log(`fccSentence.includes("freecodecamp") returns ${hasLowercaseFCC} because includes is case-sensitive.`);

const message = "Welcome to freeCodeCamp!";

console.log("Here are some examples of the slice() method:");

--fcc-editable-region--

--fcc-editable-region--
```
