---
id: 68e6172875008fea742a63c9
title: Passo 3
challengeType: 1
dashedName: step-3
---

# --description--

Como você se lembra das lições anteriores, o método `includes()` verifica se uma string contém uma substring específica e retorna `true` ou `false`.

Aqui está um exemplo usando o método `includes()`:

```js
const text = "Hello World";
const hasHello = text.includes("Hello");
console.log(hasHello); // true
```

Crie uma variável chamada `hasFreeCodeCamp`. Então atribua a ela o resultado do uso do método `includes()` em `fccSentence` para verificar se contém `"freeCodeCamp"`.

# --hints--

Você deve declarar uma variável chamada `hasFreeCodeCamp`.

```js
assert.exists(hasFreeCodeCamp);
```

Você deve usar `fccSentence.includes("freeCodeCamp")`.

```js
assert.match(__helpers.removeJSComments(code), /fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

Você deve atribuir o resultado de `fccSentence.includes("freeCodeCamp")` à sua variável `hasFreeCodeCamp`.

```js
assert.match(__helpers.removeJSComments(code), /hasFreeCodeCamp\s*=\s*fccSentence\.includes\(\s*(['"`])freeCodeCamp\1\s*\)/);
```

O resultado de `fccSentence.includes("freeCodeCamp")` deve ser `true`. 

```js
assert.isTrue(hasFreeCodeCamp);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

--fcc-editable-region--

--fcc-editable-region--
```
