---
id: 68e6172975008fea742a63ca
title: Passo 4
challengeType: 1
dashedName: step-4
---

# --description--

Agora use uma template literal ou concatenação de strings para registrar a mensagem `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.` no console. Substitua `<hasFreeCodeCamp>` pelo valor real da variável.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve registrar a mensagem `fccSentence.includes("freeCodeCamp") returns <hasFreeCodeCamp> because the word "freeCodeCamp" is in the sentence.`, onde `<hasFreeCodeCamp>` deve ser substituído pelo valor real da variável.

```js
assert.deepInclude(spy.calls, ['fccSentence.includes("freeCodeCamp") returns true because the word "freeCodeCamp" is in the sentence.']);
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

console.log("Here are some examples of the includes() method:");

const hasFreeCodeCamp = fccSentence.includes("freeCodeCamp");
--fcc-editable-region--

--fcc-editable-region--
```
