---
id: 68ef6e42fa6e8a9a18a401fa
title: Passo 2
challengeType: 1
dashedName: step-2
---

# --description--

Use `console.log()` para registrar `"Here are some examples of the includes() method:"` no console.

# --before-each--

```js
const spy = __helpers.spyOn(console, 'log');
```

# --hints--

Você deve registrar `"Here are some examples of the includes() method:"` no console.

```js
assert.deepInclude(spy.calls, ["Here are some examples of the includes() method:"])
```

# --seed--

## --seed-contents--

```js
const fccSentence = "freeCodeCamp is a great place to learn web development.";

--fcc-editable-region--

--fcc-editable-region--
```
