---
id: 694bf6eb530e19a9c48c59f3
title: 步骤 12
challengeType: 20
dashedName: step-12
---

# --description--

现在你将再次使用 `str()` 函数。就像处理 age 一样，你必须先将任何数字变量转换为字串，然后才能将其与其他文本连接。

创建一个名为 `experience_years` 的变量并赋值为整数 `5`。

然后，创建一个变量 `experience_info`。将一个由 `'Experience: '`、`experience_years` 变量（转换为字串）和 `' years'` 连接而成的字串赋值给它。将结果打印到终端。

# --hints--

你应该有一个名为 `experience_years` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_years")`)) })
```

你应该将整数 `5` 赋值给你的 `experience_years` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_years").is_equivalent("experience_years = 5")`)) })
```

你应该有一个名为 `experience_info` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("experience_info")`)) })
```

你应该将由连接 `'Experience: '`、`str(experience_years)` 和 `' years'` 组成的字串赋值给你的 `experience_info` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("experience_info").is_equivalent("experience_info = 'Experience: ' + str(experience_years) + ' years'")`)) })
```

你应该打印 `experience_info`。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(experience_info)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
--fcc-editable-region--

--fcc-editable-region--
```
