---
id: 694bf6eb530e19a9c48c59f1
title: 步骤5
challengeType: 20
dashedName: step-5
---

# --description--

现在，你的地址似乎不完整。你还想为员工居住的公寓号添加信息，所以你应该修改该变量。  

当你想将内容添加到现有字符串变量的末尾时，你可以使用 **扩充赋值** 操作符 `+=`。这比写 `var = var + 'new text'` 更简洁。例如：  

```py
greeting = 'Hello'
greeting += ' World'
print(greeting) # Hello World
```

请记住，字串是不可变的，因此此操作不会更改原始字串。相反，它会创建一个新字串并将其赋值给变量。

使用 `+=` 操作符将字串 `, Apartment 4B` 添加到你的 `address` 变量。

# --hints--

你应该使用 `+=` 操作符为你的 `address` 变量添加字串 `, Apartment 4B`。

```js
({ test: () => assert(runPython(`_Node(_code).find_aug_variable("address").is_equivalent("address += ', Apartment 4B'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
address = '123 Main Street'  
--fcc-editable-region--  

--fcc-editable-region-- 
print(address)  
```
