---
id: 694bf6eb530e19a9c48c59f7
title: 步骤18
challengeType: 20
dashedName: step-18
---

# --description--

你也可以使用负数从字串末尾切片。例如，`-1` 参考最后一个字符，`-2` 参考倒数第二个字符，依此类推。

要获取字串的最后 3 个字符，你可以使用 `string[-3:]`。注意冒号后省略了停止参数。这意味着切片应一直进行到字串边界。

创建一个名为 `last_three` 的变量。使用负索引从 `employee_code` 中提取最后三个表现 ID 号码的字符。最后，将 `last_three` 打印到终端。

至此，员工卡工作坊已完成。

# --hints--

你应该有一个名为 `last_three` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("last_three")`)) })
```

你应该使用负索引切片 `employee_code` 的最后三个字符，并将它们赋值给你的 `last_three` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_three").is_equivalent("last_three = employee_code[-3:]")`)) })
```

你应该将 `last_three` 打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_three)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
print(year_code)
initials = employee_code[9:11]
print(initials)
--fcc-editable-region--

--fcc-editable-region--
```

# --solutions--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2024-JD-001'
department = employee_code[0:3]
print(department)
year_code = employee_code[4:8]
initials = employee_code[9:11]
print(year_code)
print(initials)
last_three = employee_code[-3:]
print(last_three)
```
