---
id: 694bf6eb530e19a9c48c59f5
title: 步骤 16
challengeType: 20
dashedName: step-16
---

# --description--

在处理字串时，你经常需要提取字串的特定部分。这称为**切片**。

语法是 `string[start:stop]`，其中：

* `start` 是切片开始的索引（**包含**）。
* `stop` 是切片结束的索引（**不包含**）。

例如，如果 `text = 'Python'`，那么 `text[0:2]` 返回 `'Py'`。

将 `employee_code` 赋值为 `'DEV-2026-JD-001'`。之后，创建一个变量 `department` 并将 `employee_code` 从索引 `0` 到 `3` 的切片赋值给它。然后将 `department` 打印到终端。

# --hints--
你应该有一个名为 `employee_code` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_code")`)) })
```

你应该将字串 `DEV-2026-JD-001` 赋值给你的 `employee_code` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_code").is_equivalent("employee_code = 'DEV-2026-JD-001'")`)) })
```

你应该有一个名为 `department` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("department")`)) })
```

你应该切片 `employee_code` 的前三个字符并将它们赋值给你的 `department` 变量。

```js
({ test: () => runPython(`
dep = _Node(_code).find_variable("department")
assert dep.is_equivalent("department = employee_code[0:3]") or dep.is_equivalent("department = employee_code[:3]")
`) })
```

你应该将 `department` 打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(department)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)

--fcc-editable-region--

--fcc-editable-region--
```
