---
id: 694bf6eb530e19a9c48c59f4
title: 步骤 13
challengeType: 20
dashedName: step-13
---

# --description--

使用 `+` 连接许多字串并使用 `str()` 转换数字会变得混乱且难以读取。

Python 3.6 引入了 **f-strings** 来解决这个问题。通过在开头的引用前添加字母 `f`，你可以将变量和表达式放入由括弧状的花括弧 `{}` 表现的替换字段中。例如：

```python
name = 'John'
print(f'Hello {name}') # Output: Hello John
```

创建一个变量 `employee_card` 并赋值一个 f-字串，该字串显示 `Employee:` 后跟一个空格和 `full_name` 变量的值。

# --hints--

你应该有一个名为 `employee_card` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_card")`)) })
```

你的 `employee_card` 变量应具有值 `f'Employee: {full_name}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name}'")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

--fcc-editable-region--
```
