---
id: 694bf3af05d10fd89dab05cc
title: 步骤2
challengeType: 20
dashedName: step-2
---

# --description--

在 Python 中，你可以使用 `+` 操作符来串接字串。以下是一个示例：

```python
greeting = 'Hello' + 'World'
print(greeting) # Output: HelloWorld
```

创建一个变量 `full_name`，通过连接 `first_name` 和 `last_name`。然后打印 `full_name`。

# --hints--

你应该有一个名为 `full_name` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("full_name")`)) })
```

你应该将由连接 `first_name` 和 `last_name` 形成的字串赋值给你的 `full_name` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + last_name")`)) })
```

你应该打印 `full_name` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(full_name)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--

--fcc-editable-region--
```
