---
id: 694bf6ea530e19a9c48c59f0
title: 步骤 3
challengeType: 20
dashedName: step-3
---

# --description--

如果你将两个字串连接起来，比如 `'John' + 'Doe'`，结果是 `'JohnDoe'`，中间没有空格。要解决这个问题，你需要在它们之间连接一个包含空格的字串（`' '`）。


更新你的 `full_name` 变量，使其连接 `first_name`、一个空格和 `last_name`。

# --hints--

你应该在 `first_name` 和 `last_name` 之间添加一个空格（`' '`）。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("full_name").is_equivalent("full_name = first_name + ' ' + last_name")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
--fcc-editable-region--
full_name = first_name + last_name
--fcc-editable-region--
print(full_name)
```
