---
id: 694bf6eb530e19a9c48c59f6
title: 第 17 步
challengeType: 20
dashedName: step-17
---

# --description--

你可以从字串的任何部分切片，而不仅仅是开头。在许多分支中这都很有帮助。

创建一个变量 `year_code` 并将 `employee_code` 从索引 `4` 到 `8` 的切片赋值给它。这将提取出 `2026`。

然后创建一个变量 `initials` 并将 `employee_code` 从索引 `9` 到 `11` 的切片赋值给它。这将提取出 `JD`。

最后，将两个变量打印到终端。

# --hints--

你应该有一个名为 `year_code` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("year_code")`)) })
```

你应该从索引 `4` 到 `8` 切片 `employee_code` 并将结果赋值给你的 `year_code` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("year_code").is_equivalent("year_code = employee_code[4:8]")`)) })
```

你应该有一个名为 `initials` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("initials")`)) })
```

你应该从索引 `9` 到 `11` 切片 `employee_code` 并将结果赋值给你的 `initials` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("initials").is_equivalent("initials = employee_code[9:11]")`)) })
```

你应该将 `year_code` 打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(year_code)")`)) })
```

你应该将 `initials` 打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(initials)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
years_experience = 5
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
position = 'Data Analyst'
salary = 75000
employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'
print(employee_card)
employee_code = 'DEV-2026-JD-001'
department = employee_code[0:3]
print(department) 
--fcc-editable-region--  

--fcc-editable-region--  
```
