---
id: 695c9f21a8b1d53cf1b97841
title: 步骤7
challengeType: 20
dashedName: step-7
---

# --description--

现在创建一个名为 `employee_age` 的变量并赋值为整数 `28`。 


# --hints--

你应该有一个名为 `employee_age` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_age")`)) })
```

你应该将整数 `28` 赋值给你的 `employee_age` 变量。不要用引用将文本括起来。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_age").is_equivalent("employee_age = 28")`)) })
```


# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
--fcc-editable-region--

--fcc-editable-region--
```
