---
id: 695c9c7c61d446f1451e0ae7
title: 步骤4
challengeType: 20
dashedName: step-4
---

# --description--

接下来，创建一个变量 `address` 来保存员工的地址。为它赋值字串 `123 Main Street`，最后打印 `address`。

# --hints--

你应该有一个名为 `address` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("address")`)) })
```

你应该将字串 `123 Main Street` 赋值给你的 `address` 变量。记住要用单引号或双引号将文本括起来。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("address").is_equivalent("address = '123 Main Street'")`)) })
```

你应该打印 `address` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(address)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
print(first_name)
print(last_name)
full_name = first_name + ' ' + last_name
print(full_name)
--fcc-editable-region--

--fcc-editable-region--
```

