---
id: 695c96f82af4fb78e59fa9ed
title: 步骤 1
challengeType: 20
dashedName: step-1
---

# --description--

字串是用于保存文本数据的字符序列。正如你可能从之前的课程中回忆的那样，你可以通过将文本用单引号（`'`）或双引号（`"`）括起来来创建字串。例如：

```python
greeting = 'Hello'
print(greeting) # Output: Hello
```

创建一个变量 `first_name`，保存字串 `John`，以及一个变量 `last_name`，保存字串 `Doe`。然后打印 `first_name` 和 `last_name`。

# --hints--

你应该有一个名为 `first_name` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("first_name")`)) })
```

你应该将字串 `John` 赋值给你的 `first_name` 变量。记住要用单引号或双引号将文本括起来。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("first_name").is_equivalent("first_name = 'John'")`)) })
```

你应该有一个名为 `last_name` 的变量。

```js  
({ test: () => assert(runPython(`_Node(_code).has_variable("last_name")`)) })
```

你应该将字串 `Doe` 赋值给你的 `last_name` 变量。记住要用单引号或双引号将文本括起来。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("last_name").is_equivalent("last_name = 'Doe'")`)) })
```

你应该将你的 `first_name` 变量打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(first_name)")`)) })
```

你应该将你的 `last_name` 变量打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(last_name)")`)) })
```

# --seed--

## --seed-contents--

```python
--fcc-editable-region--

--fcc-editable-region--
```
