---
id: 695df6480de23758af5a29cd
title: 步骤 8
challengeType: 20
dashedName: step-8
---

# --description--

现在，你想创建一个显示员工年龄的字串。

首先创建一个变量 `employee_info` 并赋值为连接结果：

- `full_name` 变量。
- 一个由空格前后包围的字符 `is` 组成的字串。

# --hints--

你应该有一个名为 `employee_info` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("employee_info")`)) })
```

你应该将 `full_name + ' is '` 赋值给你的 `employee_info` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is '")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--

--fcc-editable-region--
```
