---
id: 695ea1e55c0e619e597fc5c9
title: 步骤 10
challengeType: 20
dashedName: step-10
---

# --description--

如你所见，Python 引起了一个 `TypeError: can only concatenate str (not "int") to str`。这是因为 Python 不允许你直接连接文本（字串）和数字（整数）。

为了解决这个问题，你必须先使用 `str()` 函数将数字转换为字串，该函数返回对象的字串版本：

```py
my_num = str(42)
print(type(my_num)) # <class 'str'>
```

更新你的 `employee_info` 赋值，将 `employee_age` 使用 `str(employee_age)` 转换为字串。

# --hints--

你应该通过使用 `str(employee_age)` 将 `employee_age` 转换为字串来修改你的 `employee_info` 变量的赋值。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_info").is_equivalent("employee_info = full_name + ' is ' + str(employee_age)")`)) })
```

# --seed--

## --seed-contents--

```python
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is ' + employee_age
--fcc-editable-region--
```
