---
id: 695fd3dde03273875e5d6709
title: 步骤 9
challengeType: 20
dashedName: step-9
---

# --description--

现在尝试将 `employee_age` 连接到你的 `employee_info` 字串末尾。

完成后，你将在终端看到一个 `TypeError`。在下一步中，你将着手修复它。

# --hints--

你应该将 `employee_age` 连接到 `employee_info` 的末尾。

```js
const regex = /^employee_info\s*=\s*full_name\s*\+\s*(["']) is \1\s*\+\s*employee_age\s*$/m;  
assert.match(__helpers.python.removeComments(code), regex);  
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
--fcc-editable-region--
employee_info = full_name + ' is '
--fcc-editable-region--
```
