---
id: 695fcfbba648b3a96a9495a4
title: 步骤6
challengeType: 20
dashedName: step-6
---

# --description--

终端中的输出变得拥挤了。在你继续之前，是时候清理一下了。

从你的代码中删除所有的 `print()` 语句。

# --hints--

你的代码中不应有任何 `print` 调用。

```js
({ test: () => assert(runPython(`len(_Node(_code).find_calls("print")) == 0`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'  
last_name = 'Doe'  
--fcc-editable-region--  
print(first_name)  
print(last_name)  
full_name = first_name + ' ' + last_name  
print(full_name)  
address = '123 Main Street'  
address += ', Apartment 4B'  
print(address)  
--fcc-editable-region--
```
