---
id: 69727ea5c9606bedfe36d3d2
title: 步骤15
challengeType: 20
dashedName: step-15
---

# --description--

现在是为卡片添加最终细节的时候了。

创建一个名为 `position` 的变量，值为字串 `Data Analyst`，以及一个名为 `salary` 的变量，值为整数 `75000`。

然后，更新你的 `employee_card` f-string 以包含职位和薪水。它应遵循此精确格式：`Employee: [full_name] | Age: [employee_age] | Position: [position] | Salary: $[salary]`。将占位符替换为相应的变量。

最后，打印 `employee_card` 以查看结果。

# --hints--

你应该有一个名为 `position` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).has_variable("position")`)) })
```

你应该将字串 `Data Analyst` 赋值给你的 `position` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("position").is_equivalent("position = 'Data Analyst'")`)) })
```

你应该有一个名为 `salary` 的变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

你应该将整数 `75000` 赋值给你的 `salary` 变量。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("salary").is_equivalent("salary = 75000")`)) })
```

你应该更新 `employee_card` 以使用 f-字串 `f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: ${salary}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age} | Position: {position} | Salary: \${salary}'")`)) })
```

你应该将 `employee_card` 打印到终端。

```js
({ test: () => assert(runPython(`_Node(_code).has_call("print(employee_card)")`))})
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--

employee_card = f'Employee: {full_name} | Age: {employee_age}'
--fcc-editable-region--
```
