---
id: 6980d3ecfe800d8a54130f30
title: 步骤14
challengeType: 20
dashedName: step-14
---

# --description--

目前，`employee_card` 只显示员工的姓名。现在你将为它添加更多信息。 

更新 `employee_card` 赋值以包含员工的年龄。最终的字串应如下所示：`Employee: [name] | Age: [age]`，其中 `[name]` 替换为员工的姓名，`[age]` 替换为员工的年龄。

# --hints--

你应该更新 `employee_card` 以使用 f-string `f'Employee: {full_name} | Age: {employee_age}'`。

```js
({ test: () => assert(runPython(`_Node(_code).find_variable("employee_card").is_equivalent("employee_card = f'Employee: {full_name} | Age: {employee_age}'")`)) })
```

# --seed--

## --seed-contents--

```py
first_name = 'John'
last_name = 'Doe'
full_name = first_name + ' ' + last_name
address = '123 Main Street'
address += ', Apartment 4B'
employee_age = 28
employee_info = full_name + ' is ' + str(employee_age) + ' years old'
print(employee_info)
experience_years = 5
experience_info = 'Experience: ' + str(experience_years) + ' years'
print(experience_info)
--fcc-editable-region--
employee_card = f'Employee: {full_name}'
--fcc-editable-region--
```
