---
id: 5a24c314108439a4d403614c
title: Redux 저장소에서 상태 가져오기
challengeType: 6
forumTopicId: 301443
dashedName: get-state-from-the-redux-store
---

# --description--

Redux 저장소 객체는 상호작용할 수 있는 여러 메서드를 제공합니다. 예를 들어, `state` 메서드를 사용해 Redux 저장소 객체에 저장된 현재 `getState()`를 가져올 수 있습니다.

# --instructions--

이전 도전 과제의 코드를 코드 편집기에서 더 간결하게 다시 작성했습니다. `store.getState()`를 사용해 `state`에서 `store`를 가져오고, 이를 새로운 변수 `currentState`에 할당하세요.

# --hints--

Redux 저장소의 초기 상태 값은 5여야 합니다.

```js
assert(store.getState() === 5);
```

변수 `currentState`가 존재해야 하며 Redux 저장소의 현재 상태가 할당되어야 합니다.

```js
assert(currentState === 5 && code.includes('store.getState()'));
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
```

# --solutions--

```js
const store = Redux.createStore(
  (state = 5) => state
);

// Change code below this line
const currentState = store.getState();
```
