---
id: 5a24c314108439a4d4036152
title: 액션 타입에 const 사용하기
challengeType: 6
forumTopicId: 301450
dashedName: use-const-for-action-types
---

# --description--

Redux를 사용할 때 흔히 하는 관례는 액션 타입을 읽기 전용 상수로 할당하고, 이 상수를 사용하는 모든 곳에서 참조하는 것입니다. 작업 중인 코드를 리팩터링하여 액션 타입을 `const` 선언으로 작성하세요.

# --instructions--

`LOGIN`과 `LOGOUT`를 `const` 값으로 선언하고 각각 문자열 `'LOGIN'`와 `'LOGOUT'`를 할당하세요. 그런 다음 `authReducer()`과 액션 생성자들이 문자열 값 대신 이 상수를 참조하도록 수정하세요.

**참고:** 상수는 일반적으로 모두 대문자로 작성하는 관례가 있으며, Redux에서도 표준 관례입니다.

# --hints--

함수 `loginUser`을 호출하면 `type` 속성이 문자열 `LOGIN`로 설정된 객체가 반환되어야 합니다.

```js
assert(loginUser().type === 'LOGIN');
```

함수 `logoutUser`을 호출하면 `type` 속성이 문자열 `LOGOUT`로 설정된 객체가 반환되어야 합니다.

```js
assert(logoutUser().type === 'LOGOUT');
```

저장소는 `login` 속성이 `false`로 설정된 객체로 초기화되어야 합니다.

```js
assert(store.getState().authenticated === false);
```

`loginUser`를 디스패치하면 저장소 상태의 `login` 속성이 `true`로 업데이트되어야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

`logoutUser`를 디스패치하면 저장소 상태의 `login` 속성이 `false`로 업데이트되어야 합니다.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

`authReducer` 함수는 switch 문으로 여러 액션 타입을 처리해야 합니다.

```js
assert(
    (function () {
      return (
        typeof authReducer === 'function' &&
        code.toString().includes('switch') &&
        code.toString().includes('case') &&
        code.toString().includes('default')
      );
    })()
  );
```

`LOGIN`와 `LOGOUT`는 `const` 값으로 선언되어야 하며, 각각 `LOGIN`와 `LOGOUT` 문자열이 할당되어야 합니다.

```js
const noWhiteSpace = __helpers.removeWhiteSpace(code);
assert(LOGIN === 'LOGIN' && LOGOUT === 'LOGOUT')
assert(noWhiteSpace.includes('const'))
```

액션 생성자와 리듀서는 `LOGIN`와 `LOGOUT` 상수를 참조해야 합니다.

```js
  assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(
        code.toString()
      );
      return (
        noWhiteSpace.includes('caseLOGIN:') &&
        noWhiteSpace.includes('caseLOGOUT:') &&
        noWhiteSpace.includes('type:LOGIN') &&
        noWhiteSpace.includes('type:LOGOUT')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js


const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {
    case 'LOGIN': 
      return {
        authenticated: true
      }
    case 'LOGOUT': 
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case LOGIN:
      return {
        authenticated: true
      }

    case LOGOUT:
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: LOGIN
  }
};

const logoutUser = () => {
  return {
    type: LOGOUT
  }
};
```
