---
id: 5a24c314108439a4d403614d
title: Redux 액션 정의하기
challengeType: 6
forumTopicId: 301440
dashedName: define-a-redux-action
---

# --description--

Redux는 상태 관리 프레임워크이므로 상태를 업데이트하는 것이 핵심 작업 중 하나입니다. Redux에서는 모든 상태 업데이트가 액션을 디스패치하는 것으로 시작됩니다. 액션은 발생한 이벤트에 대한 정보를 담고 있는 단순한 JavaScript 객체입니다. Redux 저장소는 이 액션 객체를 받아 상태를 그에 맞게 업데이트합니다. 때로는 Redux 액션이 데이터를 함께 전달하기도 합니다. 예를 들어, 사용자가 로그인한 후 액션이 사용자 이름을 전달할 수 있습니다. 데이터는 선택 사항이지만, 액션은 반드시 발생한 액션의 'type'을 지정하는 `type` 속성을 포함해야 합니다.

Redux 액션을 앱에서 발생하는 이벤트에 대한 정보를 Redux 저장소에 전달하는 메신저로 생각하세요. 저장소는 전달받은 액션을 바탕으로 상태를 업데이트합니다.

# --instructions--

Redux 액션을 작성하는 것은 type 속성을 가진 객체를 선언하는 것만큼 간단합니다. `action` 객체를 선언하고 `type` 속성을 문자열 `'LOGIN'`로 설정하세요.

# --hints--

`action` 객체가 존재해야 합니다.

```js
assert(
  (function () {
    return typeof action === 'object';
  })()
);
```

`action` 객체는 키 `type`를 가지고 값은 `LOGIN`이어야 합니다.

```js
assert(
  (function () {
    return action.type === 'LOGIN';
  })()
);
```

# --seed--

## --seed-contents--

```js
// Define an action here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
```
