---
id: 5a24c314108439a4d4036154
title: 여러 리듀서 결합하기
challengeType: 6
forumTopicId: 301436
dashedName: combine-multiple-reducers
---

# --description--

애플리케이션 상태가 점점 복잡해지면 상태를 여러 조각으로 나누고 싶을 수 있습니다. 하지만 Redux의 첫 번째 원칙을 기억하세요: 모든 애플리케이션 상태는 저장소의 단일 상태 객체에 보관됩니다. 따라서 Redux는 복잡한 상태 모델을 위한 해결책으로 리듀서 합성을 제공합니다. 애플리케이션 상태의 서로 다른 조각을 처리하는 여러 리듀서를 정의한 다음, 이 리듀서들을 하나의 루트 리듀서로 합성합니다. 루트 리듀서는 Redux `createStore()` 메서드에 전달됩니다.

여러 리듀서를 함께 결합할 수 있도록 Redux는 `combineReducers()` 메서드를 제공합니다. 이 메서드는 객체를 인수로 받아, 특정 리듀서 함수에 키를 연결하는 속성을 정의합니다. 키에 부여한 이름은 Redux가 해당 상태 조각의 이름으로 사용합니다.

일반적으로 애플리케이션 상태가 구별되거나 고유한 경우, 각 상태 조각마다 리듀서를 만드는 것이 좋은 습관입니다. 예를 들어, 사용자 인증이 있는 노트 작성 애플리케이션에서는 한 리듀서가 인증을 처리하고, 다른 리듀서는 사용자가 제출하는 텍스트와 노트를 처리할 수 있습니다. 이런 애플리케이션에서는 `combineReducers()` 메서드를 다음과 같이 작성할 수 있습니다:

```js
const rootReducer = Redux.combineReducers({
  auth: authenticationReducer,
  notes: notesReducer
});
```

이제 `notes` 키는 노트와 관련된 모든 상태를 포함하며 `notesReducer`가 처리합니다. 이렇게 여러 리듀서를 합성하여 더 복잡한 애플리케이션 상태를 관리할 수 있습니다. 이 예제에서 Redux 저장소에 보관된 상태는 `auth`와 `notes` 속성을 포함하는 단일 객체가 됩니다.

# --instructions--

코드 편집기에는 `counterReducer()`와 `authReducer()` 함수와 Redux 저장소가 제공됩니다. `rootReducer()` 메서드를 사용하여 `Redux.combineReducers()` 함수를 완성하세요. `counterReducer`을 `count`라는 키에, `authReducer`을 `auth`라는 키에 할당하세요.

# --hints--

`counterReducer`는 `state`를 증가 및 감소시켜야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState().count;
    store.dispatch({ type: INCREMENT });
    store.dispatch({ type: INCREMENT });
    const firstState = store.getState().count;
    store.dispatch({ type: DECREMENT });
    const secondState = store.getState().count;
    return firstState === initialState + 2 && secondState === firstState - 1;
  })()
);
```

`authReducer`는 `state`의 `authenticated`를 `true`와 `false` 사이에서 토글해야 합니다.

```js
assert(
  (function () {
    store.dispatch({ type: LOGIN });
    const loggedIn = store.getState().auth.authenticated;
    store.dispatch({ type: LOGOUT });
    const loggedOut = store.getState().auth.authenticated;
    return loggedIn === true && loggedOut === false;
  })()
);
```

저장소 `state`는 두 개의 키를 가져야 합니다: 숫자를 보관하는 `count`와 객체를 보관하는 `auth`. `auth` 객체는 불리언을 보관하는 `authenticated` 속성을 가져야 합니다.

```js
assert(
  (function () {
    const state = store.getState();
    return (
      typeof state.auth === 'object' &&
      typeof state.auth.authenticated === 'boolean' &&
      typeof state.count === 'number'
    );
  })()
);
```

`rootReducer`는 `counterReducer`와 `authReducer`를 결합하는 함수여야 합니다.

```js
assert(
    (function () {
      const noWhiteSpace = __helpers.removeWhiteSpace(code);
      return (
        typeof rootReducer === 'function' &&
        noWhiteSpace.includes('Redux.combineReducers')
      );
    })()
  );
```

# --seed--

## --seed-contents--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = // Define the root reducer here

const store = Redux.createStore(rootReducer);
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const LOGIN = 'LOGIN';
const LOGOUT = 'LOGOUT';

const authReducer = (state = {authenticated: false}, action) => {
  switch(action.type) {
    case LOGIN:
      return {
        authenticated: true
      }
    case LOGOUT:
      return {
        authenticated: false
      }
    default:
      return state;
  }
};

const rootReducer = Redux.combineReducers({
  count: counterReducer,
  auth: authReducer
});

const store = Redux.createStore(rootReducer);
```
