---
id: 5a24c314108439a4d4036155
title: 저장소에 액션 데이터를 보내기
challengeType: 6
forumTopicId: 301448
dashedName: send-action-data-to-the-store
---

# --description--

지금까지 Redux 저장소에 액션을 디스패치하는 방법을 배웠지만, 지금까지 이 액션들은 `type` 외에는 어떤 정보도 포함하지 않았습니다. 액션과 함께 특정 데이터를 보낼 수도 있습니다. 사실, 액션은 보통 사용자 상호작용에서 시작되며 데이터를 함께 전달하는 경우가 많기 때문에 매우 흔한 일입니다. Redux 저장소는 이 데이터를 알아야 할 때가 많습니다.

# --instructions--

코드 편집기에는 기본 `notesReducer()`와 `addNoteText()` 액션 생성자가 정의되어 있습니다. `addNoteText()` 함수의 본문을 완성하여 `action` 객체를 반환하게 하세요. 이 객체는 `type` 속성을 `ADD_NOTE` 값으로 포함해야 하며, 또한 액션 생성자에 전달된 `text` 데이터를 `note` 속성으로 설정해야 합니다. 액션 생성자를 호출할 때 특정 노트 정보를 전달하며, 이 정보를 객체에서 접근할 수 있습니다.

다음으로 `switch` 안의 `notesReducer()` 문장을 완성하세요. `addNoteText()` 액션을 처리하는 케이스를 추가해야 합니다. 이 케이스는 `ADD_NOTE` 타입의 액션이 있을 때 실행되며, 들어오는 `text`의 `action` 속성을 새로운 `state`로 반환해야 합니다.

코드 하단에서 액션이 디스패치됩니다. 완료하면 코드를 실행하고 콘솔을 확인하세요. 액션별 데이터를 저장소에 보내고 저장소 `state`을 업데이트할 때 사용하는 방법은 이게 전부입니다.

# --hints--

액션 생성자 `addNoteText`는 `type`와 `text` 키를 가진 객체를 반환해야 합니다.

```js
assert(
  (function () {
    const addNoteFn = addNoteText('__TEST__NOTE');
    return addNoteFn.type === ADD_NOTE && addNoteFn.text === '__TEST__NOTE';
  })()
);
```

`ADD_NOTE` 타입의 액션을 `addNoteText` 액션 생성자로 디스패치하면, `state`가 액션 생성자에 전달된 문자열로 업데이트되어야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(addNoteText('__TEST__NOTE'));
    const newState = store.getState();
    return initialState !== newState && newState === '__TEST__NOTE';
  })()
);
```

# --seed--

## --seed-contents--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line

    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello!'));
console.log(store.getState());
```

# --solutions--

```js
const ADD_NOTE = 'ADD_NOTE';

const notesReducer = (state = 'Initial State', action) => {
  switch(action.type) {
    // Change code below this line
    case ADD_NOTE:
      return action.text;
    // Change code above this line
    default:
      return state;
  }
};

const addNoteText = (note) => {
  // Change code below this line
  return {
    type: ADD_NOTE,
    text: note
  }
  // Change code above this line
};

const store = Redux.createStore(notesReducer);

console.log(store.getState());
store.dispatch(addNoteText('Hello Redux!'));
console.log(store.getState());
```
