---
id: 5a24c314108439a4d403614e
title: 액션 생성자 정의하기
challengeType: 6
forumTopicId: 301441
dashedName: define-an-action-creator
---

# --description--

액션을 만든 후 다음 단계는 Redux 저장소에 액션을 보내 상태를 업데이트하는 것입니다. Redux에서는 이를 위해 액션 생성자를 정의합니다. 액션 생성자는 단순히 액션을 반환하는 JavaScript 함수입니다. 다시 말해, 액션 생성자는 액션 이벤트를 나타내는 객체를 만듭니다.

# --instructions--

호출 시 `actionCreator()` 객체를 반환하는 `action`라는 이름의 함수를 정의하세요.

# --hints--

`actionCreator` 함수가 존재해야 합니다.

```js
assert(typeof actionCreator === 'function');
```

`actionCreator` 함수를 실행하면 `action` 객체가 반환되어야 합니다.

```js
assert(typeof action === 'object');
```

반환된 `action`에는 키 속성 `type`가 값 `LOGIN`으로 있어야 합니다.

```js
assert(action.type === 'LOGIN');
```

# --seed--

## --seed-contents--

```js
const action = {
  type: 'LOGIN'
}
// Define an action creator here:
```

# --solutions--

```js
const action = {
  type: 'LOGIN'
}
const actionCreator = () => {
  return action;
};
```
