---
id: 5a24c314108439a4d403614b
title: Redux 저장소 만들기
challengeType: 6
forumTopicId: 301439
dashedName: create-a-redux-store
---

# --description--

Redux는 React를 포함한 여러 웹 기술과 함께 사용할 수 있는 상태 관리 프레임워크입니다.

Redux에서는 애플리케이션 전체 상태를 담당하는 단일 상태 객체가 있습니다. 즉, 만약 10개의 컴포넌트가 있는 React 애플리케이션이 있고 각 컴포넌트가 자체 로컬 상태를 가지고 있다면, 애플리케이션 전체 상태는 Redux `store`에 저장된 단일 상태 객체로 정의됩니다. 이것이 Redux를 배울 때 이해해야 할 첫 번째 중요한 원칙입니다: Redux 저장소는 애플리케이션 상태에 관한 단일 진실의 출처입니다.

또한 애플리케이션의 어떤 부분이든 상태를 업데이트하려면 반드시 Redux 저장소를 통해서만 해야 합니다. 단방향 데이터 흐름 덕분에 애플리케이션의 상태 관리를 추적하기가 더 쉬워집니다.

# --instructions--

Redux `store`는 애플리케이션 `state`를 보관하고 관리하는 객체입니다. Redux 객체에는 `createStore()`라는 메서드가 있는데, 이 메서드를 사용해 Redux `store`를 만듭니다. 이 메서드는 `reducer` 함수 하나를 필수 인수로 받습니다. `reducer` 함수는 이후 도전 과제에서 다루며, 코드 편집기에 이미 정의되어 있습니다. 이 함수는 `state`를 인수로 받아 `state`를 반환합니다.

`store` 변수를 선언하고 `createStore()` 메서드에 `reducer`를 인수로 전달하여 할당하세요.

**참고:** 편집기 내 코드는 ES6 기본 인수 구문을 사용해 이 상태를 `5` 값으로 초기화합니다. 기본 인수에 익숙하지 않다면 <a href="https://www.freecodecamp.org/learn/javascript-algorithms-and-data-structures/es6/set-default-parameters-for-your-functions" target="_blank" rel="noopener noreferrer nofollow">커리큘럼의 ES6 섹션</a>을 참고하세요.

# --hints--

Redux 저장소가 존재해야 합니다.

```js
assert(typeof store.getState === 'function');
```

Redux 저장소의 상태 값은 5여야 합니다.

```js
assert(store.getState() === 5);
```

# --seed--

## --seed-contents--

```js
const reducer = (state = 5) => {
  return state;
}

// Redux methods are available from a Redux object
// For example: Redux.createStore()
// Define the store here:
```

# --solutions--

```js
const reducer = (state = 5) => {
  return state;
}

const store = Redux.createStore(reducer);
```
