---
id: 5a24c314108439a4d403615b
title: Object.assign으로 객체 복사하기
challengeType: 6
forumTopicId: 301437
dashedName: copy-an-object-with-object-assign
---

# --description--

최근 도전 과제들은 배열을 다뤘지만, 상태가 `object`일 때 상태 불변성을 유지하는 데 도움이 되는 방법도 있습니다. 객체를 다룰 때 유용한 도구가 `Object.assign()` 유틸리티입니다. `Object.assign()`는 대상 객체와 소스 객체들을 받아 소스 객체들의 속성을 대상 객체에 매핑합니다. 일치하는 속성은 소스 객체의 속성으로 덮어씌워집니다. 이 동작은 첫 번째 인수로 빈 객체를 전달하고 복사하려는 객체를 뒤따라 전달하여 객체의 얕은 복사본을 만드는 데 자주 사용됩니다. 예를 들어 다음과 같습니다:

```js
const newObject = Object.assign({}, obj1, obj2);
```

이 코드는 `newObject`를 새로운 `object`로 생성하는데, 여기에는 현재 `obj1`와 `obj2`에 존재하는 속성들이 포함됩니다.

# --instructions--

Redux 상태와 액션이 `object`를 `state`에 처리하도록 수정되었습니다. 타입이 `state`인 액션에 대해 `ONLINE` 속성을 문자열 `status`로 설정하는 새로운 `online` 객체를 반환하도록 코드를 수정하세요. 도전 과제를 완료하려면 `Object.assign()`를 사용해 보세요.

# --hints--

Redux 저장소는 존재해야 하며 1번째 줄에 선언된 `defaultState` 객체와 동등한 상태로 초기화되어야 합니다.

```js
const expectedState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};
const initialState = store.getState();
assert.deepEqual(expectedState, initialState);
```

`wakeUp`와 `immutableReducer`는 모두 함수여야 합니다.

```js
assert(typeof wakeUp === 'function' && typeof immutableReducer === 'function');
```

타입이 `ONLINE`인 액션을 디스패치하면 상태의 `status` 속성이 `online`로 업데이트되어야 하며 상태를 변경해서는 안 됩니다.

```js

const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch({ type: 'ONLINE' });
const finalState = store.getState();
const expectedState = {
  user: 'CamperBot',
  status: 'online',
  friends: '732,982',
  community: 'freeCodeCamp'
};
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

새 상태를 반환할 때 `Object.assign`를 사용해야 합니다.

```js
assert(code.includes('Object.assign'));
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const defaultState = {
  user: 'CamperBot',
  status: 'offline',
  friends: '732,982',
  community: 'freeCodeCamp'
};

const immutableReducer = (state = defaultState, action) => {
  switch(action.type) {
    case 'ONLINE':
      return Object.assign({}, state, {
        status: 'online'
      });
    default:
      return state;
  }
};

const wakeUp = () => {
  return {
    type: 'ONLINE'
  }
};

const store = Redux.createStore(immutableReducer);
```
