---
id: 5a24c314108439a4d4036157
title: Redux로 카운터 만들기
challengeType: 6
forumTopicId: 301453
dashedName: write-a-counter-with-redux
---

# --description--

이제 Redux의 핵심 원칙을 모두 배웠습니다! 액션과 액션 생성자를 만드는 방법, Redux 저장소를 생성하는 방법, 저장소에 액션을 디스패치하는 방법, 순수한 리듀서로 상태 업데이트를 설계하는 방법을 보았습니다. 리듀서 조합으로 복잡한 상태를 관리하고 비동기 액션을 처리하는 방법도 살펴봤습니다. 이 예제들은 단순하지만, 이 개념들이 Redux의 핵심 원칙입니다. 이 원칙들을 잘 이해했다면 자신만의 Redux 앱을 만들 준비가 된 것입니다. 다음 도전 과제에서는 `state` 불변성에 관한 세부 사항을 다루지만, 먼저 지금까지 배운 내용을 검토해 보겠습니다.

# --instructions--

이번 수업에서는 Redux로 간단한 카운터를 처음부터 구현합니다. 기본 코드는 코드 편집기에 제공되지만 세부 사항을 채워야 합니다! 제공된 이름을 사용하고 `incAction`과 `decAction` 액션 생성자, `counterReducer()`, `INCREMENT`, `DECREMENT` 액션 타입, 그리고 마지막으로 Redux `store`를 정의하세요. 완료하면 `INCREMENT` 또는 `DECREMENT` 액션을 디스패치하여 `store`에 저장된 상태를 증가시키거나 감소시킬 수 있어야 합니다. 첫 Redux 앱을 만드는 데 행운을 빕니다!

# --hints--

액션 생성자 `incAction`는 `type`이 `INCREMENT` 값과 같은 액션 객체를 반환해야 합니다.

```js
assert(incAction().type === INCREMENT);
```

액션 생성자 `decAction`는 `type`이 `DECREMENT` 값과 같은 액션 객체를 반환해야 합니다.

```js
assert(decAction().type === DECREMENT);
```

`store.getState()`를 실행하면 숫자가 반환되어야 합니다.

```js
assert(typeof store.getState() === 'number');
```

Redux 저장소는 `state`이 0으로 초기화되어야 합니다.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    return _store.getState() === 0;
  })()
);
```

Redux 저장소에 `incAction`를 디스패치하면 `state`가 1만큼 증가해야 합니다.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(incAction());
    const incState = _store.getState();
    return initialState + 1 === incState;
  })()
);
```

Redux 저장소에 `decAction`를 디스패치하면 `state`가 1만큼 감소해야 합니다.

```js
assert(
  (function () {
    const _store = Redux.createStore(counterReducer);
    const initialState = _store.getState();
    _store.dispatch(decAction());
    const decState = _store.getState();
    return initialState - 1 === decState;
  })()
);
```

`counterReducer`는 함수여야 합니다.

```js
assert(typeof counterReducer === 'function');
```

# --seed--

## --seed-contents--

```js
const INCREMENT = null; // Define a constant for increment action types
const DECREMENT = null; // Define a constant for decrement action types

const counterReducer = null; // Define the counter reducer which will increment or decrement the state based on the action it receives

const incAction = null; // Define an action creator for incrementing

const decAction = null; // Define an action creator for decrementing

const store = null; // Define the Redux store here, passing in your reducers
```

# --solutions--

```js
const INCREMENT = 'INCREMENT';
const DECREMENT = 'DECREMENT';

const counterReducer = (state = 0, action) => {
  switch(action.type) {
    case INCREMENT:
      return state + 1;
    case DECREMENT:
      return state - 1;
    default:
      return state;
  }
};

const incAction = () => {
  return {
    type: INCREMENT
  }
};

const decAction = () => {
  return {
    type: DECREMENT
  }
};

const store = Redux.createStore(counterReducer);
```
