---
id: 5a24c314108439a4d403614f
title: 액션 이벤트 디스패치하기
challengeType: 6
forumTopicId: 301442
dashedName: dispatch-an-action-event
---

# --description--

`dispatch` 메서드는 Redux 저장소에 액션을 디스패치할 때 사용합니다. `store.dispatch()`을 호출하고 액션 생성자에서 반환된 값을 전달하면 액션이 저장소로 전송됩니다.

액션 생성자가 발생한 액션 유형을 지정하는 type 속성을 가진 객체를 반환한다는 점을 기억하세요. 그런 다음 이 메서드는 액션 객체를 Redux 저장소에 디스패치합니다. 이전 도전 과제 예제를 기준으로 다음 두 줄은 동일하며, 둘 다 `LOGIN` 유형의 액션을 디스패치합니다.

```js
store.dispatch(actionCreator());
store.dispatch({ type: 'LOGIN' });
```

# --instructions--

코드 편집기의 Redux 저장소는 현재 `login`로 설정된 `false` 속성을 포함하는 객체로 초기화된 상태를 가지고 있습니다. `loginAction()` 유형의 액션을 반환하는 `LOGIN`라는 액션 생성자도 있습니다. `LOGIN` 메서드를 호출하여 `dispatch` 액션을 Redux 저장소에 디스패치하고, `loginAction()`가 생성한 액션을 전달하세요.

# --hints--

`loginAction` 함수를 호출하면 `type` 속성이 `LOGIN` 문자열로 설정된 객체가 반환되어야 합니다.

```js
assert(loginAction().type === 'LOGIN');
```

저장소는 `login` 속성이 `false`로 설정된 객체로 초기화되어야 합니다.

```js
assert(store.getState().login === false);
```

`store.dispatch()` 메서드를 사용하여 `LOGIN` 유형의 액션을 디스패치해야 합니다.

```js
assert(
    (function () {
      let noWhiteSpace = code.replace(/\s/g, '');
      return (
        noWhiteSpace.includes('store.dispatch(loginAction())') ||
        noWhiteSpace.includes("store.dispatch({type: 'LOGIN'})") === true
      );
    })()
);
```

# --seed--

## --seed-contents--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

// Dispatch the action here:
```

# --solutions--

```js
const store = Redux.createStore(
  (state = {login: false}) => state
);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};

store.dispatch(loginAction());
```
