---
id: 5a24c314108439a4d4036159
title: 배열에 전개 연산자 사용하기
challengeType: 6
forumTopicId: 301452
dashedName: use-the-spread-operator-on-arrays
---

# --description--

Redux에서 상태 불변성을 강제하는 데 도움이 되는 ES6의 한 가지 해결책은 전개 연산자입니다: `...`. 전개 연산자는 다양한 용도로 사용되며, 그중 하나는 기존 배열에서 새 배열을 만드는 이전 도전 과제에 적합합니다. 이 구문은 비교적 새롭지만 널리 사용됩니다. 예를 들어, 배열 `myArray`이 있고 다음과 같이 작성하면:

```js
let newArray = [...myArray];
```

`newArray`는 이제 `myArray`의 복제본입니다. 두 배열은 메모리에서 별도로 존재합니다. `newArray.push(5)`와 같은 변이를 수행해도 `myArray`는 변경되지 않습니다. `...`는 `myArray`의 값을 새 배열로 *펼치는* 역할을 합니다. 배열을 복제하되 새 배열에 추가 값을 넣으려면 `[...myArray, 'new value']`와 같이 작성할 수 있습니다. 이렇게 하면 `myArray`의 값과 문자열 `new value`이 마지막 값으로 포함된 새 배열이 반환됩니다. 전개 구문은 배열 구성에서 여러 번 사용할 수 있지만, 배열을 얕게 복사한다는 점을 유의하세요. 즉, 1차원 배열에 대해서만 불변 배열 연산을 제공합니다.

# --instructions--

할 일이 추가될 때 전개 연산자를 사용해 상태의 새 복사본을 반환하세요.

# --hints--

Redux 저장소는 존재해야 하며 상태가 `["Do not mutate state!"]`과 같도록 초기화해야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    return (
      Array.isArray(initialState) === true &&
      initialState[0] === 'Do not mutate state!'
    );
  })()
);
```

`addToDo`와 `immutableReducer`는 모두 함수여야 합니다.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Redux 저장소에서 `ADD_TO_DO` 유형의 액션을 디스패치하면 `todo` 항목이 추가되어야 하며 상태가 변이되어서는 안 됩니다.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = ['Do not mutate state!', '__TEST__TO__DO__'];
assert(isFrozen)
assert.deepEqual(finalState, expectedState);
```

새 상태를 반환할 때 전개 연산자를 사용해야 합니다.

```js
assert(code.includes('...state'));
```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = ['Do not mutate state!'], action) => {
  switch(action.type) {
    case 'ADD_TO_DO':
      return [
        ...state,
        action.todo
      ];
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: 'ADD_TO_DO',
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
