---
id: 5a24c314108439a4d4036150
title: 저장소에서 액션 처리하기
challengeType: 6
forumTopicId: 301444
dashedName: handle-an-action-in-the-store
---

# --description--

액션이 생성되어 디스패치된 후, Redux 저장소는 그 액션에 어떻게 응답할지 알아야 합니다. 이것이 `reducer` 함수의 역할입니다. Redux의 리듀서는 액션에 대한 응답으로 상태를 수정하는 일을 담당합니다. `reducer`은 `state`와 `action`를 인수로 받고, 항상 새로운 `state`를 반환합니다. 이것이 리듀서의 **유일한** 역할임을 이해하는 것이 중요합니다. 리듀서는 부작용이 없으며 — API 엔드포인트를 호출하지 않고 숨겨진 동작도 없습니다. 리듀서는 단순히 상태와 액션을 받아 새로운 상태를 반환하는 순수한 함수입니다.

Redux의 또 다른 핵심 원칙은 `state`가 읽기 전용이라는 점입니다. 즉, `reducer` 함수는 **항상** `state`의 새 복사본을 반환해야 하며 상태를 직접 수정해서는 안 됩니다. Redux가 상태 불변성을 강제하지는 않지만, 리듀서 함수 코드에서 이를 지키는 것은 작성자의 책임입니다. 이후 도전 과제에서 이를 연습할 것입니다.

# --instructions--

코드 편집기에는 이전 예제와 `reducer` 함수의 시작 부분이 있습니다. `reducer` 함수의 본문을 채워서, 만약 `'LOGIN'` 타입의 액션을 받으면 `login`이 `true`로 설정된 상태 객체를 반환하게 하세요. 그렇지 않으면 현재 `state`를 반환합니다. 현재 `state`와 디스패치된 `action`가 리듀서에 전달되므로, `action.type`를 통해 액션의 타입에 직접 접근할 수 있습니다.

# --hints--

`loginAction` 함수를 호출하면 문자열 `LOGIN`가 설정된 type 속성을 가진 객체를 반환해야 합니다.

```js
assert(loginAction().type === 'LOGIN');
```

저장소는 `login` 속성이 `false`로 설정된 객체로 초기화되어야 합니다.

```js
assert(store.getState().login === false);
```

`loginAction`를 디스패치하면 저장소 상태의 `login` 속성이 `true`로 업데이트되어야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginAction());
    const afterState = store.getState();
    return initialState.login === false && afterState.login === true;
  })()
);
```

액션이 `LOGIN` 타입이 아니면 저장소는 현재 상태를 반환해야 합니다.

```js
assert(
  (function () {
    store.dispatch({ type: '__TEST__ACTION__' });
    let afterTest = store.getState();
    return typeof afterTest === 'object' && afterTest.hasOwnProperty('login');
  })()
);
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```

# --solutions--

```js
const defaultState = {
  login: false
};

const reducer = (state = defaultState, action) => {

  if (action.type === 'LOGIN') {
    return {login: true}
  }

  else {
    return state
  }

};

const store = Redux.createStore(reducer);

const loginAction = () => {
  return {
    type: 'LOGIN'
  }
};
```
