---
id: 5a24c314108439a4d4036153
title: 저장소 리스너 등록
challengeType: 6
forumTopicId: 301446
dashedName: register-a-store-listener
---

# --description--

Redux `store` 객체에서 접근할 수 있는 또 다른 메서드는 `store.subscribe()`입니다. 이 메서드는 저장소에 리스너 함수를 구독하게 하며, 저장소에 액션이 디스패치될 때마다 호출됩니다. 이 메서드의 간단한 용도 중 하나는 액션이 수신되고 저장소가 업데이트될 때마다 메시지를 기록하는 함수를 저장소에 구독시키는 것입니다.

# --instructions--

저장소가 액션을 받을 때마다 전역 변수 `count`를 증가시키는 콜백 함수를 작성하고, 이 함수를 `store.subscribe()` 메서드에 전달하세요. `store.dispatch()`가 연속으로 세 번 호출되며, 매번 액션 객체를 직접 전달하는 것을 볼 수 있습니다. 액션 디스패치 사이의 콘솔 출력을 확인하여 업데이트가 이루어지는 것을 확인하세요.

# --hints--

저장소에서 `ADD` 액션을 디스패치하면 상태가 `1`만큼 증가해야 합니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch({ type: 'ADD' });
    const newState = store.getState();
    return newState === initialState + 1;
  })()
);
```

저장소에 `store.subscribe`를 사용해 리스너 함수가 구독되어 있어야 합니다.

```js
assert.match(code, /store\s*\.\s*subscribe\(/gm);
```

`store.subscribe`는 함수를 받아야 합니다.

```js
assert.match(code, /(\s*function\s*)|(\s*\(\s*\)\s*=>)/gm);
```

`store.subscribe`에 전달된 함수는 호출되어서는 안 됩니다.

```js
assert.notMatch(code, /store\.subscribe\(.+\(\)\)/);
```

`store.subscribe`에 대한 콜백은 저장소가 업데이트될 때 전역 `count` 변수도 증가시켜야 합니다.

```js
assert.strictEqual(store.getState(), count);
```

# --seed--

## --seed-contents--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);

// Global count variable:
let count = 0;

// Change code below this line

// Change code above this line

store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
store.dispatch({type: ADD});
console.log(count);
```

# --solutions--

```js
const ADD = 'ADD';

const reducer = (state = 0, action) => {
  switch(action.type) {
    case ADD:
      return state + 1;
    default:
      return state;
  }
};

const store = Redux.createStore(reducer);
 let count = 0;
// Change code below this line

store.subscribe( () =>
 {
 count++;
 }
);

// Change code above this line

store.dispatch({type: ADD});
store.dispatch({type: ADD});
store.dispatch({type: ADD});
```
