---
id: 5a24c314108439a4d403615a
title: 배열에서 항목 제거하기
challengeType: 6
forumTopicId: 301447
dashedName: remove-an-item-from-an-array
---

# --description--

배열에서 항목을 제거하는 연습을 할 시간입니다. 전개 연산자도 여기서 사용할 수 있습니다. 다른 유용한 JavaScript 메서드로는 `slice()`와 `concat()`가 있습니다.

# --instructions--

리듀서와 액션 생성자가 항목의 인덱스를 기준으로 배열에서 항목을 제거하도록 수정되었습니다. 특정 인덱스의 항목이 제거된 새 상태 배열이 반환되도록 리듀서 작성을 마무리하세요.

# --hints--

Redux 저장소는 존재해야 하며 상태는 `[0,1,2,3,4,5]`와 같아야 합니다.

```js
const initialState = store.getState();
assert.isArray(initialState) 
assert.deepEqual(initialState, [0, 1, 2, 3, 4, 5])
```

`removeItem`와 `immutableReducer`는 모두 함수여야 합니다.

```js
assert(
  typeof removeItem === 'function' && typeof immutableReducer === 'function'
);
```

`removeItem` 액션 생성자를 디스패치하면 상태에서 항목이 제거되어야 하며 상태를 변경해서는 안 됩니다.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(removeItem(3));
const state_1 = store.getState();
store.dispatch(removeItem(2));
const state_2 = store.getState();
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
store.dispatch(removeItem(0));
const state_3 = store.getState();

assert(isFrozen)
assert.deepEqual(state_1, [0, 1, 2, 4, 5]) 
assert.deepEqual(state_2, [0, 1, 4, 5]) 
assert.deepEqual(state_3, [5])

```

# --seed--

## --seed-contents--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const immutableReducer = (state = [0,1,2,3,4,5], action) => {
  switch(action.type) {
    case 'REMOVE_ITEM':
      return [
        ...state.slice(0, action.index),
        ...state.slice(action.index + 1)
      ];
    default:
      return state;
  }
};

const removeItem = (index) => {
  return {
    type: 'REMOVE_ITEM',
    index
  }
}

const store = Redux.createStore(immutableReducer);
```
