---
id: 5a24c314108439a4d4036151
title: 여러 액션을 처리하기 위해 switch 문 사용하기
challengeType: 6
forumTopicId: 301449
dashedName: use-a-switch-statement-to-handle-multiple-actions
---

# --description--

Redux 저장소에 여러 액션 유형을 처리하는 방법을 알려줄 수 있습니다. 예를 들어 Redux 저장소에서 사용자 인증을 관리한다고 가정해 보세요. 사용자가 로그인했을 때와 로그아웃했을 때의 상태를 나타내는 상태 표현이 필요합니다. 이 상태는 `authenticated` 속성을 가진 단일 상태 객체로 나타냅니다. 또한 사용자 로그인과 로그아웃에 해당하는 액션을 생성하는 액션 생성자와 액션 객체도 필요합니다.

# --instructions--

코드 편집기에는 저장소, 액션, 액션 생성자가 설정되어 있습니다. 여러 인증 액션을 처리하도록 `reducer` 함수를 작성하세요. JavaScript `switch` 문을 `reducer` 안에서 사용하여 다양한 액션 이벤트에 대응하세요. 이것은 Redux 리듀서를 작성할 때 표준 패턴입니다. switch 문은 `action.type`를 기준으로 분기하며 적절한 인증 상태를 반환해야 합니다.

**참고:** 이 시점에서는 상태 불변성에 대해 걱정하지 마세요. 이 예제는 작고 단순하기 때문입니다. 각 액션에 대해 새 객체를 반환할 수 있습니다 — 예를 들어 `{authenticated: true}`. 또한 switch 문에 현재 `default`을 반환하는 `state` 케이스를 작성하는 것을 잊지 마세요. 이는 앱에 여러 리듀서가 있을 때 모든 액션 디스패치 시 모든 리듀서가 실행되기 때문에 중요합니다. 액션이 해당 리듀서와 관련이 없을 때는 현재 `state`을 반환해야 합니다.

# --hints--

`loginUser` 함수를 호출하면 type 속성이 문자열 `LOGIN`으로 설정된 객체가 반환되어야 합니다.

```js
assert(loginUser().type === 'LOGIN');
```

`logoutUser` 함수를 호출하면 type 속성이 문자열 `LOGOUT`으로 설정된 객체가 반환되어야 합니다.

```js
assert(logoutUser().type === 'LOGOUT');
```

저장소는 `authenticated` 속성이 `false`로 설정된 객체로 초기화되어야 합니다.

```js
assert(store.getState().authenticated === false);
```

`loginUser`를 디스패치하면 저장소 상태의 `authenticated` 속성이 `true`로 업데이트됩니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(loginUser());
    const afterLogin = store.getState();
    return (
      initialState.authenticated === false && afterLogin.authenticated === true
    );
  })()
);
```

`logoutUser`를 디스패치하면 저장소 상태의 `authenticated` 속성이 `false`로 업데이트됩니다.

```js
assert(
  (function () {
    store.dispatch(loginUser());
    const loggedIn = store.getState();
    store.dispatch(logoutUser());
    const afterLogout = store.getState();
    return (
      loggedIn.authenticated === true && afterLogout.authenticated === false
    );
  })()
);
```

`authReducer` 함수는 `switch` 문으로 여러 액션 유형을 처리해야 합니다.

```js
assert(
    code.toString().includes('switch') &&
      code.toString().includes('case') &&
      code.toString().includes('default')
  );
```

# --seed--

## --seed-contents--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {
  // Change code below this line

  // Change code above this line
};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```

# --solutions--

```js
const defaultState = {
  authenticated: false
};

const authReducer = (state = defaultState, action) => {

  switch (action.type) {

    case 'LOGIN':
      return {
        authenticated: true
      }

    case 'LOGOUT':
      return {
        authenticated: false
      }

    default:
      return state;

  }

};

const store = Redux.createStore(authReducer);

const loginUser = () => {
  return {
    type: 'LOGIN'
  }
};

const logoutUser = () => {
  return {
    type: 'LOGOUT'
  }
};
```
