---
id: 5a24c314108439a4d4036158
title: 상태를 절대 변경하지 마세요
challengeType: 6
forumTopicId: 301445
dashedName: never-mutate-state
---

# --description--

이 마지막 도전 과제들은 Redux에서 상태 불변성이라는 핵심 원칙을 지키는 여러 방법을 설명합니다. 불변 상태란 상태를 직접 수정하지 않고, 대신 상태의 새 복사본을 반환하는 것을 의미합니다.

Redux 앱의 상태를 시간에 따라 스냅샷으로 찍으면 `state 1`, `state 2`, `state 3`, `state 4`, `...` 같은 모습이 보일 것입니다. 각 상태는 이전 상태와 비슷할 수 있지만, 각각은 별개의 데이터 조각입니다. 이 불변성이 바로 여러분이 들어봤을 수도 있는 타임트래블 디버깅 같은 기능을 가능하게 합니다.

Redux는 저장소나 리듀서에서 상태 불변성을 적극적으로 강제하지 않으며, 그 책임은 프로그래머에게 있습니다. 다행히도 JavaScript(특히 ES6)는 상태가 `string`, `number`, `array` 또는 `object`이든 상관없이 상태 불변성을 지키는 데 유용한 여러 도구를 제공합니다. 문자열과 숫자는 원시 값이며 본질적으로 불변임을 유의하세요. 즉, 3은 항상 3이며, 숫자 3의 값을 변경할 수 없습니다. 그러나 `array` 또는 `object`은 변경 가능합니다. 실제로 상태는 여러 유형의 정보를 나타내는 데 유용한 자료 구조인 `array` 또는 `object`로 구성될 가능성이 큽니다.

# --instructions--

코드 편집기에는 할 일 항목을 관리하는 `store`과 `reducer`가 있습니다. 리듀서에서 `ADD_TO_DO` 케이스를 완성하여 상태에 새 할 일을 추가하세요. 표준 JavaScript나 ES6로 이를 수행하는 몇 가지 방법이 있습니다. `action.todo`의 항목을 끝에 추가한 새 배열을 반환하는 방법을 찾아보세요.

# --hints--

Redux 저장소는 존재해야 하며, 코드 편집기의 `todos` 배열과 같은 상태로 초기화되어야 합니다.

```js
assert(
  (function () {
    const todos = [
      'Go to the store',
      'Clean the house',
      'Cook dinner',
      'Learn to code'
    ];
    const initialState = store.getState();
    return (
      Array.isArray(initialState) && initialState.join(',') === todos.join(',')
    );
  })()
);
```

`addToDo`과 `immutableReducer`는 모두 함수여야 합니다.

```js
assert(typeof addToDo === 'function' && typeof immutableReducer === 'function');
```

Redux 저장소에서 `ADD_TO_DO` 타입의 액션을 디스패치하면 `todo` 항목이 추가되어야 하며 상태를 변경해서는 안 됩니다.

```js
const initialState = store.getState();
const isFrozen = !!DeepFreeze(initialState);
store.dispatch(addToDo('__TEST__TO__DO__'));
const finalState = store.getState();
const expectedState = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
  '__TEST__TO__DO__'
];
assert(isFrozen);
assert.deepEqual(finalState, expectedState);
```

# --seed--

## --seed-contents--

```js
const ADD_TO_DO = 'ADD_TO_DO';

// A list of strings representing tasks to do:
const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      // Don't mutate state here or the tests will fail
      return
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```

# --solutions--

```js
const ADD_TO_DO = 'ADD_TO_DO';

const todos = [
  'Go to the store',
  'Clean the house',
  'Cook dinner',
  'Learn to code',
];

const immutableReducer = (state = todos, action) => {
  switch(action.type) {
    case ADD_TO_DO:
      return state.concat(action.todo);
    default:
      return state;
  }
};

const addToDo = (todo) => {
  return {
    type: ADD_TO_DO,
    todo
  }
}

const store = Redux.createStore(immutableReducer);
```
