---
id: 5a24c314108439a4d4036156
title: 비동기 액션을 처리하기 위해 미들웨어 사용하기
challengeType: 6
forumTopicId: 301451
dashedName: use-middleware-to-handle-asynchronous-actions
---

# --description--

지금까지의 도전 과제에서는 비동기 액션에 대해 다루지 않았지만, 비동기 액션은 웹 개발에서 피할 수 없는 부분입니다. 언젠가는 Redux 애플리케이션에서 비동기 엔드포인트를 호출해야 할 텐데, 이런 유형의 요청을 어떻게 처리할까요? Redux는 이를 위해 특별히 설계된 미들웨어인 Redux Thunk 미들웨어를 제공합니다. Redux와 함께 사용하는 방법을 간단히 설명합니다.

Redux Thunk 미들웨어를 포함하려면 `Redux.applyMiddleware()`에 인수로 전달하세요. 이 문장은 `createStore()` 함수에 두 번째 선택적 매개변수로 제공됩니다. 편집기 하단의 코드를 확인해 보세요. 그런 다음 비동기 액션을 만들려면 액션 생성자에서 `dispatch`를 인수로 받는 함수를 반환하세요. 이 함수 내에서 액션을 디스패치하고 비동기 요청을 수행할 수 있습니다.

이 예제에서는 `setTimeout()` 호출로 비동기 요청을 시뮬레이션합니다. 비동기 동작을 시작하기 전에 액션을 디스패치하는 것이 일반적이며, 이렇게 하면 애플리케이션 상태가 데이터를 요청 중임을 알 수 있습니다(예를 들어 로딩 아이콘을 표시할 수 있습니다). 그런 다음 데이터를 받으면, 이 데이터를 페이로드로 포함하고 액션이 완료되었음을 알리는 또 다른 액션을 디스패치합니다.

이 특별한 액션 생성자에 `dispatch`를 매개변수로 전달한다는 점을 기억하세요. 이 매개변수를 사용해 액션을 디스패치하며, 액션을 dispatch에 직접 전달하면 미들웨어가 나머지를 처리합니다.

# --instructions--

`handleAsync()` 액션 생성자 안에 두 개의 디스패치를 작성하세요. `requestingData()` (시뮬레이션된 API 호출) 전에 `setTimeout()`를 디스패치하세요. 그런 다음 (가상의) 데이터를 받은 후에는 이 데이터를 전달하며 `receivedData()` 액션을 디스패치하세요. 이제 Redux에서 비동기 액션을 처리하는 방법을 알게 되었습니다. 나머지는 이전과 동일하게 작동합니다.

# --hints--

`requestingData` 액션 생성자는 `REQUESTING_DATA` 값과 같은 타입의 객체를 반환해야 합니다.

```js
assert(requestingData().type === REQUESTING_DATA);
```

`receivedData` 액션 생성자는 `RECEIVED_DATA` 값과 같은 타입의 객체를 반환해야 합니다.

```js
assert(receivedData('data').type === RECEIVED_DATA);
```

`asyncDataReducer`는 함수여야 합니다.

```js
assert(typeof asyncDataReducer === 'function');
```

`requestingData` 액션 생성자를 디스패치하면 저장소의 `state` 속성이 `true`로 업데이트됩니다.

```js
assert(
  (function () {
    const initialState = store.getState();
    store.dispatch(requestingData());
    const reqState = store.getState();
    return initialState.fetching === false && reqState.fetching === true;
  })()
);
```

`handleAsync`를 디스패치하면 데이터 요청 액션을 디스패치한 후 지연 후에 받은 데이터 액션을 디스패치합니다.

```js
assert(
  (function () {
    const noWhiteSpace = __helpers.removeWhiteSpace(handleAsync.toString());
    return (
      noWhiteSpace.includes('dispatch(requestingData())') === true &&
      noWhiteSpace.includes('dispatch(receivedData(data))') === true
    );
  })()
);
```

# --seed--

## --seed-contents--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    // Dispatch request action here

    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      // Dispatch received data action here

    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```

# --solutions--

```js
const REQUESTING_DATA = 'REQUESTING_DATA'
const RECEIVED_DATA = 'RECEIVED_DATA'

const requestingData = () => { return {type: REQUESTING_DATA} }
const receivedData = (data) => { return {type: RECEIVED_DATA, users: data.users} }

const handleAsync = () => {
  return function(dispatch) {
    dispatch(requestingData());
    setTimeout(function() {
      let data = {
        users: ['Jeff', 'William', 'Alice']
      }
      dispatch(receivedData(data));
    }, 2500);
  }
};

const defaultState = {
  fetching: false,
  users: []
};

const asyncDataReducer = (state = defaultState, action) => {
  switch(action.type) {
    case REQUESTING_DATA:
      return {
        fetching: true,
        users: []
      }
    case RECEIVED_DATA:
      return {
        fetching: false,
        users: action.users
      }
    default:
      return state;
  }
};

const store = Redux.createStore(
  asyncDataReducer,
  Redux.applyMiddleware(ReduxThunk.default)
);
```
