---
id: 668f08ea07b99b1f4a91acab
title: 創建一個食譜頁面
challengeType: 25
dashedName: build-a-recipe-page
demoType: onClick
---

# --description--

**目標：**實現以下用戶需求並通過所有測試以完成實驗。

**用戶需求：**

1. 你應該有一個 `!DOCTYPE html` 聲明。
1. 你應該有一個 `lang` 設置爲 `en` 的 `html` 元素。
1. 你應該有一個 `head` 元素，其中包含一個用於你的食譜名稱的 `title` 元素和一個 `charset` 屬性設置爲 `UTF-8` 的 `meta` 元素。
1. 你應該有一個 `body` 元素。
1. 你應該有一個包含你的食譜名稱的 `h1` 元素。
1. 在 `h1` 下面，你應該有一個介紹食譜的 `p` 元素。
1. 你應該有一個文本爲 `Ingredients` 的 `h2`元素，用於原料部分。
1. 在第一個 `h2` 元素下面，你應該有一個至少有四個列表項（`li` 元素）的無序列表（`ul` 元素），列出你的原料。
1. 你應該一個有文本爲 `Instructions` 的第二個 `h2` 元素，用於操作步驟部分。
1. 在第二個 `h2` 元素下面，你應該有一個至少有四個列表項的有序列表（`ol` 元素），按順序列出食譜步驟。
1. 你應該有一個 `img` 元素，它有一個設置爲有效圖像的 `src` 屬性，如果你願意可以使用 `https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg`，和一個描述圖像的 `alt` 屬性。

# --hints--

你的食譜頁面該有一個 `!DOCTYPE html` 聲明。

```js
assert.match(code, /<!DOCTYPE html>/i);
```

你應該有一個 `lang` 設置爲 `en` 的 `html` 元素。

```js
assert.match(code, /<html\s+lang\s*=\s*('|")en\1\s*>[\s\S]*<\/\s*html\s*>/gi);
```

你的 `html` 元素中應該有一個 `head` 元素。

```js
assert.match(code, /<html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\/\s*html\s*>/i);
```

你的 `head` 元素中應該有一個 `title` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*title\s*>[\s\S]*<\/\s*title\s*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `title` 元素應該包含你的食譜標題。

```js
assert.isAbove(document.querySelector('title')?.innerText.trim().length, 0);
```

你的 `head` 元素中應該有一個 `meta` 元素。

```js
assert.match(code, /<\s*head\s*>[\s\S]*<\s*meta[\s\S]*>[\s\S]*<\/\s*head\s*>/i);
```

你的 `meta` 元素應該將它的 `charset` 屬性設置爲 `UTF-8`。

```js
assert.match(code, /<\s*meta[\s\S]+?charset\s*=\s*('|")UTF-8\1/i);
```

你的 `html` 元素中應該有一個 `body` 元素。

```js
assert.match(code, /<\s*html[\s\S]*>[\s\S]*<\s*head\s*>[\s\S]*<\/\s*head\s*>[\s\S]*<\s*body\s*>[\s\S]*<\/\s*body\s*>[\s\S]*<\/\s*html\s*>/i);
```

你應該有一個包含你的食譜名稱的 `h1` 元素。

```js
assert.isAbove(document.querySelector('h1')?.innerText.length, 0);
```

You should only have one `h1` element.

```js
assert.lengthOf(document.querySelectorAll('h1'), 1);
```

你的 `h1` 元素下面應該有一個 `p` 元素。

```js
assert.strictEqual(document.querySelector('h1')?.nextElementSibling, document.querySelector('p'));
```

你的第一個 `p` 元素應該描述你的食譜。

```js
assert.isNotEmpty(document.querySelector('p')?.textContent?.trim());
```

你的第一個 `h2` 元素應該包含文本 `Ingredients`。

```js
assert.equal(document.querySelectorAll('h2')[0]?.innerText, 'Ingredients');
```

你的第一個 `h2` 元素下面應該有一個無序列表元素。

```js
assert.strictEqual(document.querySelector('ul')?.previousElementSibling.tagName, 'H2');
```

你的無序列表中應該至少有四個包含原料的列表項元素。

```js
const els = document.querySelectorAll('ul > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你的第二個 `h2` 元素應該包含文本 `Instructions`。

```js
assert.equal(document.querySelectorAll('h2')[1]?.innerText, 'Instructions');
```

你的第二個 `h2` 元素下面應該有一個有序列表元素。

```js
assert.strictEqual(document.querySelectorAll('h2')?.[1]?.nextElementSibling?.tagName, "OL");
```

你的無序列表中應該至少有四個包含操作步驟的列表項元素。

```js
const els = document.querySelectorAll('ol > li');
assert.isAbove(els.length, 3);
els.forEach(el => assert.isAbove(el.innerText.trim().length, 0))
```

你應該至少有一個 `img` 元素。

```js
assert.exists(document.querySelector('img'));
```

你所有的 `img` 元素都應該有一個有效的 `src` 屬性和值。

```js
const img = document.querySelector('img');
const rawSrc = img?.getAttribute('src');
const resolvedSrc = img?.src;
const re = new RegExp(window.location.href, "ig");

assert.isAbove(rawSrc?.trim().length, 0, "The 'src' attribute must be explicitly set.");
assert.notMatch(resolvedSrc, re, "The 'src' should not start with the current page URL");

img.onload = () => {
  console.log('Image loaded successfully.');
};

img.onerror = (error) => {
  console.error('Image failed to load:', error);
  assert.fail("Your image's URL should be valid."); // Make the test instafail
};

if (img.complete) {
  img.onload && img.onload();
};
```

你所有的 `img` 元素都應該有一個 `alt` 屬性來描述圖像。

```js
assert.isAbove(document.querySelector('img')?.alt.length, 0);
```

# --seed--

## --seed-contents--

```html

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
  <meta charset="utf-8">
  <title>Chocolate chip cookies recipe</title>
</head>

<body>
  <h1>Chocolate Chip Cookies</h1>
  <p>Welcome to the ultimate guide for making mini chocolate chip cookies! These bite-sized treats are perfect for
    satisfying your sweet tooth without overindulging. Follow this simple recipe to create delicious,
    crispy-on-the-outside, chewy-on-the-inside mini chocolate chip cookies that everyone will love.</p>
  <img src="https://cdn.freecodecamp.org/curriculum/labs/recipe.jpg" alt="Ingredients for baking: three eggs, a bowl of flour, a glass of milk, and a whisk arranged on a wooden table.">
  <h2>Ingredients</h2>
  <ul>
    <li>1 cup all-purpose flour</li>
    <li>1/2 teaspoon baking soda</li>
    <li>1/4 cup unsalted butter, softened</li>
    <li>1/4 cup granulated sugar</li>
    <li>1/2 teaspoon vanilla extract</li>
    <li>1/2 cup mini chocolate chips</li>
  </ul>
  <h2>Instructions</h2>
  <ol>
    <li>Preheat your oven to 350°F (175°C) and line a baking sheet with parchment paper.</li>
    <li>In a bowl, whisk together the flour and baking soda.</li>
    <li>In another bowl, beat the butter, sugar, and vanilla extract until creamy.</li>
    <li>Gradually add the dry ingredients to the wet mixture, then fold in the mini chocolate chips.</li>
    <li>Drop small spoonfuls of dough onto the baking sheet.</li>
    <li>Bake for 8-10 minutes, then let cool before enjoying!</li>
  </ol>
</body>

</html>
```
