---
id: 5a24c314108439a4d4036143
title: Statuslogik in Redux auslagern
challengeType: 6
forumTopicId: 301429
dashedName: extract-state-logic-to-redux
---

# --description--

Nachdem Sie die React-Komponente fertiggestellt haben, müssen Sie die Logik, die lokal in ihrem `state` ausgeführt wird, in Redux auslagern. Dies ist der erste Schritt, um die einfache React-App mit Redux zu verbinden. Die einzige Funktionalität Ihrer App besteht darin, neue Nachrichten vom Benutzer zu einer ungeordneten Liste hinzuzufügen. Das Beispiel ist einfach gehalten, um zu demonstrieren, wie React und Redux zusammenarbeiten.

# --instructions--

Definieren Sie zunächst einen Action-Typ `ADD` und setzen Sie ihn auf eine Konstante `ADD`. Definieren Sie anschließend einen Action Creator `addMessage()`, der die Aktion zum Hinzufügen einer Nachricht erstellt. Sie müssen diesem Action Creator einen `message` übergeben und die Nachricht im zurückgegebenen `action` einschließen.

Erstellen Sie dann einen Reducer namens `messageReducer()`, der den Status für die Nachrichten verwaltet. Der Anfangszustand sollte ein leeres Array sein. Dieser Reducer soll eine Nachricht zum Array der im Status gehaltenen Nachrichten hinzufügen oder den aktuellen Status zurückgeben. Erstellen Sie abschließend Ihren Redux-Store und übergeben Sie ihm den Reducer.

# --hints--

Die Konstante `ADD` sollte existieren und einen Wert haben, der dem String `ADD` entspricht.

```js
assert(ADD === 'ADD');
```

Der Action Creator `addMessage` sollte ein Objekt zurückgeben, bei dem `type` gleich `ADD` und `message` gleich der übergebenen Nachricht ist.

```js
assert(
  (function () {
    const addAction = addMessage('__TEST__MESSAGE__');
    return addAction.type === ADD && addAction.message === '__TEST__MESSAGE__';
  })()
);
```

`messageReducer` sollte eine Funktion sein.

```js
assert(typeof messageReducer === 'function');
```

Der Store sollte existieren und einen Anfangszustand haben, der ein leeres Array ist.

```js
assert(
  (function () {
    const initialState = store.getState();
    return typeof store === 'object' && initialState.length === 0;
  })()
);
```

Das Dispatchen von `addMessage` gegen den Store sollte eine neue Nachricht unveränderlich zum Array der im Status gehaltenen Nachrichten hinzufügen.

```js
assert(
  (function () {
    const initialState = store.getState();
    const isFrozen = DeepFreeze(initialState);
    store.dispatch(addMessage('__A__TEST__MESSAGE'));
    const addState = store.getState();
    return isFrozen && addState[0] === '__A__TEST__MESSAGE';
  })()
);
```

`messageReducer` sollte den aktuellen Status zurückgeben, wenn es mit anderen Aktionen aufgerufen wird.

```js
assert(
  (function () {
    const addState = store.getState();
    store.dispatch({ type: 'FAKE_ACTION' });
    const testState = store.getState();
    return addState === testState;
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Define ADD, addMessage(), messageReducer(), and store here:
```

# --solutions--

```jsx
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);
```
