---
id: 5a24c314108439a4d4036147
title: Redux mit React verbinden
challengeType: 6
forumTopicId: 301426
dashedName: connect-redux-to-react
---

# --description--

Jetzt, da Sie sowohl die `mapStateToProps()`- als auch die `mapDispatchToProps()`-Funktionen geschrieben haben, können Sie diese verwenden, um `state` und `dispatch` auf die `props` einer Ihrer React-Komponenten abzubilden. Die `connect`-Methode von React Redux kann diese Aufgabe übernehmen. Diese Methode nimmt zwei optionale Argumente, `mapStateToProps()` und `mapDispatchToProps()`. Sie sind optional, weil Sie möglicherweise eine Komponente haben, die nur Zugriff auf `state` benötigt, aber keine Aktionen dispatchen muss, oder umgekehrt.

Um diese Methode zu verwenden, übergeben Sie die Funktionen als Argumente und rufen das Ergebnis sofort mit Ihrer Komponente auf. Diese Syntax ist etwas ungewöhnlich und sieht so aus:

```js
connect(mapStateToProps, mapDispatchToProps)(MyComponent)
```

**Hinweis:** Wenn Sie eines der Argumente der `connect`-Methode weglassen möchten, übergeben Sie stattdessen `null`.

# --instructions--

Der Code-Editor enthält die `mapStateToProps()`- und `mapDispatchToProps()`-Funktionen sowie eine neue React-Komponente namens `Presentational`. Verbinden Sie diese Komponente mit Redux über die `connect`-Methode des globalen `ReactRedux`-Objekts und rufen Sie sie sofort mit der `Presentational`-Komponente auf. Weisen Sie das Ergebnis einer neuen `const` namens `ConnectedComponent` zu, die die verbundene Komponente darstellt. Das war's, jetzt sind Sie mit Redux verbunden! Versuchen Sie, eines der Argumente von `connect` in `null` zu ändern, und beobachten Sie die Testergebnisse.

# --hints--

Die `Presentational`-Komponente sollte gerendert werden.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('Presentational').length === 1;
  })()
);
```

Die `Presentational`-Komponente sollte eine Prop `messages` über `connect` erhalten.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return props.messages === '__INITIAL__STATE__';
  })()
);
```

Die `Presentational`-Komponente sollte eine Prop `submitNewMessage` über `connect` erhalten.

```js
assert(
  (function () {
    if (typeof ConnectedComponent === 'undefined') {
      return false;
    }
    const store = Redux.createStore(
      (state = '__INITIAL__STATE__', action) => state
    );
    class AppWrapper extends React.Component {
      render() {
        return React.createElement(
          ReactRedux.Provider,
          { store: store },
          React.createElement(ConnectedComponent)
        );
      }
    }
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    const props = mockedComponent.find('Presentational').props();
    return typeof props.submitNewMessage === 'function';
  })()
);
```

# --seed--

## --seed-contents--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line
```

# --solutions--

```jsx
const addMessage = (message) => {
  return {
    type: 'ADD',
    message: message
  }
};

const mapStateToProps = (state) => {
  return {
    messages: state
  }
};

const mapDispatchToProps = (dispatch) => {
  return {
    submitNewMessage: (message) => {
      dispatch(addMessage(message));
    }
  }
};

class Presentational extends React.Component {
  constructor(props) {
    super(props);
  }
  render() {
    return <h3>This is a Presentational Component</h3>
  }
};

const connect = ReactRedux.connect;
// Change code below this line

const ConnectedComponent = connect(mapStateToProps, mapDispatchToProps)(Presentational);
```
