---
id: 5a24c314108439a4d4036144
title: "Provider verwenden, um Redux mit React zu verbinden"
challengeType: 6
forumTopicId: 301435
dashedName: use-provider-to-connect-redux-to-react
---

# --description--

Im letzten Schritt haben Sie einen Redux-Store erstellt, um das Array der Nachrichten zu verwalten, und eine Aktion zum Hinzufügen neuer Nachrichten definiert. Der nächste Schritt besteht darin, React den Zugriff auf den Redux-Store und die Aktionen zu ermöglichen, die es zum Auslösen von Updates benötigt. React Redux stellt sein `react-redux`-Paket bereit, um diese Aufgaben zu erleichtern.

React Redux bietet eine kleine API mit zwei Hauptfunktionen: `Provider` und `connect`. Eine weitere Aufgabe behandelt `connect`. Der `Provider` ist eine Wrapper-Komponente von React Redux, die Ihre React-App umschließt. Dieser Wrapper ermöglicht es Ihnen dann, im gesamten Komponentenbaum auf den Redux-`store` und die `dispatch`-Funktionen zuzugreifen. `Provider` nimmt zwei Props entgegen, den Redux-Store und die Kind-Komponenten Ihrer App. Die Definition des `Provider` für eine App-Komponente könnte so aussehen:

```jsx
<Provider store={store}>
  <App/>
</Provider>
```

# --instructions--

Der Code-Editor zeigt jetzt all Ihren Redux- und React-Code aus den letzten Aufgaben. Er enthält den Redux-Store, Aktionen und die `DisplayMessages`-Komponente. Das einzige neue Element ist die `AppWrapper`-Komponente am Ende. Verwenden Sie diese oberste Komponente, um den `Provider` von `ReactRedux` zu rendern und übergeben Sie den Redux-Store als Prop. Dann rendern Sie die `DisplayMessages`-Komponente als Kind. Sobald Sie fertig sind, sollten Sie Ihre React-Komponente auf der Seite sehen.

**Hinweis:** React Redux ist hier als globale Variable verfügbar, sodass Sie mit Punktnotation auf den Provider zugreifen können. Der Code im Editor nutzt dies und weist ihn einer Konstanten `Provider` zu, die Sie in der `AppWrapper`-Render-Methode verwenden können.

# --hints--

Der `AppWrapper` sollte gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return mockedComponent.find('AppWrapper').length === 1;
  })()
);
```

Die `Provider`-Wrapper-Komponente sollte eine Prop `store` erhalten, die dem Redux-Store entspricht.

```js
assert(
    (function () {
      const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
      return __helpers
        .removeWhiteSpace(code)
        .includes('<Providerstore={store}>');
    })()
  );
```

`DisplayMessages` sollte als Kind von `AppWrapper` gerendert werden.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('AppWrapper').find('DisplayMessages').length === 1
    );
  })()
);
```

Die `DisplayMessages`-Komponente sollte ein `h2`-, `input`-, `button`- und `ul`-Element rendern.

```js
assert(
  (function () {
    const mockedComponent = Enzyme.mount(React.createElement(AppWrapper));
    return (
      mockedComponent.find('div').length === 1 &&
      mockedComponent.find('h2').length === 1 &&
      mockedComponent.find('button').length === 1 &&
      mockedComponent.find('ul').length === 1
    );
  })()
);
```

# --seed--

## --seed-contents--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};



const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
    this.handleChange = this.handleChange.bind(this);
    this.submitMessage = this.submitMessage.bind(this);
  }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {  
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Render the Provider below this line

  // Change code above this line
};
```

# --solutions--

```jsx
// Redux:
const ADD = 'ADD';

const addMessage = (message) => {
  return {
    type: ADD,
    message
  }
};

const messageReducer = (state = [], action) => {
  switch (action.type) {
    case ADD:
      return [
        ...state,
        action.message
      ];
    default:
      return state;
  }
};

const store = Redux.createStore(messageReducer);

// React:

class DisplayMessages extends React.Component {
  constructor(props) {
    super(props);
    this.state = {
      input: '',
      messages: []
    }
 this.handleChange = this.handleChange.bind(this);
 this.submitMessage = this.submitMessage.bind(this);
 }
  handleChange(event) {
    this.setState({
      input: event.target.value
    });
  }
  submitMessage() {
    this.setState((state) => {
      const currentMessage = state.input;
      return {
        input: '',
        messages: state.messages.concat(currentMessage)
      };  
    });
  }
  render() {
    return (
      <div>
        <h2>Type in a new Message:</h2>
        <input
          value={this.state.input}
          onChange={this.handleChange}/><br/>
        <button onClick={this.submitMessage}>Submit</button>
        <ul>
          {this.state.messages.map( (message, idx) => {
              return (
                 <li key={idx}>{message}</li>
              )
            })
          }
        </ul>
      </div>
    );
  }
};

const Provider = ReactRedux.Provider;

class AppWrapper extends React.Component {
  // Change code below this line
  render() {
    return (
      <Provider store = {store}>
        <DisplayMessages/>
      </Provider>
    );
  }
  // Change code above this line
};
```
