---
id: 66f53dc2c5bd6a11d6c3282f
title: Construir um Aplicativo Gerador de Senhas
challengeType: 26
dashedName: lab-password-generator
---

# --description--

Neste laboratório, você praticará o uso de funções construindo um gerador de senhas aleatórias.

**Objetivo:** Cumprir as user stories abaixo e fazer todos os testes passarem para completar o laboratório.

**Histórias de Usuário:**

1. Você deve criar uma função chamada `generatePassword` que recebe um parâmetro, indicando o comprimento da senha gerada. Você pode nomear o parâmetro como preferir.
2. Sua função deve retornar uma string que represente uma senha gerada aleatoriamente. Você deve usar a seguinte string e diferentes métodos do `Math` para ajudá-lo a retornar uma nova string com caracteres aleatórios nela: `ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()`.
3. Você deve definir uma variável chamada `password` e atribuir a ela o resultado da chamada da função `generatePassword` com um argumento numérico que representa o comprimento desejado da senha.
4. Você deve ter um `console.log` que registre uma única string feita pela concatenação da mensagem `Generated password:` e da variável `password` separadas por um espaço.

# --hints--

Você deve ter uma função `generatePassword` com um parâmetro. Você pode nomear o parâmetro como quiser.

```js
assert.isFunction(generatePassword);
assert.lengthOf(generatePassword, 1);
```

Sua função `generatePassword` deve retornar uma string.

```js
const result = generatePassword(5);
assert.isString(result);
```

Sua função `generatePassword` deve retornar uma nova string que tenha o comprimento correto.

```js
const length = 8;
const password = generatePassword(length);
assert.lengthOf(password, length);
```

Sua função deve retornar uma senha gerada aleatoriamente com caracteres válidos.

```js
const validChars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
const password = generatePassword(10);
for (let char of password) {
    assert.include(validChars, char);
}
```

Sua função deve retornar uma senha gerada aleatoriamente que contenha mais de um caractere único.

```js
const passwordSetSize1 = new Set(generatePassword(12)).size;
const passwordSetSize2 = new Set(generatePassword(12)).size;
const passwordSetSize3 = new Set(generatePassword(12)).size;
assert.isAbove(passwordSetSize1 + passwordSetSize2 + passwordSetSize3, 3);
```

Sua função deve retornar uma nova string aleatória cada vez que for chamada.

```js
const password1 = generatePassword(10);
const password2 = generatePassword(10);
assert.notStrictEqual(password1, password2);
```

Você deve ter uma variável `password`.

```js
assert.isDefined(password);
```

Sua variável `password` deve ser uma string.

```js
assert.isString(password);
```

Você deve chamar a função `generatePassword` com um argumento numérico e armazenar a senha retornada na variável `password`.

```js
assert.isDefined(password);
assert.isString(password);
assert.match(password, /^[A-Za-z0-9!@#$%^&*()]+$/);
const length2 = 10;
const password3 = generatePassword(length2);
assert.lengthOf(password3, length2);
assert.equal(password3.length, generatePassword(length2).length);
assert.match(__helpers.removeJSComments(code), /(let|const|var)\s+password\s*=\s*generatePassword\(\d+\)\;?/);
```

Você deve registrar uma única string combinando `Senha gerada:` e a `senha` separadas por um único espaço usando `+` ou um template literal.

```js
const cleanCode = __helpers.removeJSComments(code);
const condition1 = /console\.log\(\s*["']Generated\s+password:\s*["']\s*\+\s*password\s*\);?/gm.test(cleanCode);
const condition2 = /console\.log\(\s*`Generated\s+password:\s*\$\{password\}`\s*\);?/gm.test(cleanCode);
assert.isTrue(condition1 || condition2);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const generatePassword = (length) => {
    const chars = "ABCDEFGHIJKLMNOPQRSTUVWXYZabcdefghijklmnopqrstuvwxyz0123456789!@#$%^&*()";
    let password = "";
    for (let i = 0; i < length; i++) {
        const randomIndex = Math.floor(Math.random() * chars.length);
        password += chars[randomIndex];
    }
    return password;
};

const password = generatePassword(12);
console.log(`Generated password: ${password}`);
```
