---
id: 68f1f993460f06a64c033061
title: 7단계
challengeType: 0
dashedName: step-7
---

# --description--

`poster` 속성은 비디오의 썸네일 이미지입니다. YouTube에서 보는 동영상을 생각해 보세요. 동영상이 다운로드되는 동안 표시됩니다. 속성이 생략되면 다운로드 중에 첫 번째 비디오 프레임이 표시됩니다.

이제 `poster` 요소에 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg` 속성을 값 `video`로 추가하세요.

# --hints--

`video` 요소에 `poster` 속성을 값 `https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg`로 추가해야 합니다.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('poster'), 'https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
--fcc-editable-region--
    
--fcc-editable-region--
  >
  </video>
</body>
</html>
```
