---
id: 68f1f371fefd448b4f70b898
title: 3단계
challengeType: 0
dashedName: step-3
---

# --description--

이전 수업에서 `video` 요소에 사용할 수 있는 다양한 속성에 대해 배웠습니다. `width` 속성은 비디오의 너비를 픽셀 단위로 결정합니다.

`width` 요소에 `video` 속성을 값 `640`로 추가하세요.

# --hints--

`video` 요소에 `width` 속성을 값 `640`로 추가해야 합니다.

```js
const video = document.querySelector('video');
assert.strictEqual(video?.getAttribute('width'), '640');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video>
--fcc-editable-region--
  </video>
</body>
</html>
```
