---
id: 68f1fac259a4afabb478b960
title: 8단계
challengeType: 0
dashedName: step-8
---

# --description--

실제 비디오에 링크하지 않은 것을 눈치챘을 수 있습니다. 다음 단계에서 그렇게 할 것입니다. 비디오 파일 형식에는 브라우저 지원 차이가 있습니다. 이를 해결하려면 `source` 요소 안에 `video` 요소들을 넣으면 브라우저가 첫 번째 호환 가능한 `source`를 선택합니다.

다음은 `source` 요소의 예입니다:

```html
<video controls width="250">
  <source src="src-url-goes-here" type="video-type-goes-here" />
</video>
```

`source` 요소는 빈 요소이므로 닫는 태그가 없습니다.

`source` 요소 안에 `video` 요소를 추가하세요.


# --hints--

`source` 요소 안에 `video` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('video > source'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
