---
id: 68f1f6f97af9749a42260a43
title: 5단계
challengeType: 0
dashedName: step-5
---

# --description--

`controls` 속성은 `video` 요소에 재생, 되감기, 볼륨 조절을 포함한 재생 컨트롤을 제공합니다.

`controls` 속성은 불리언 속성으로 값이 필요 없습니다.

`controls` 요소에 `video` 속성을 추가하세요.

이제 페이지에 `video` 요소가 표시되는 것을 볼 수 있습니다.

# --hints--

`video` 요소에 `controls` 속성이 있어야 합니다.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('controls'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop>
--fcc-editable-region--
  </video>
</body>
</html>
```
