---
id: 68f1f57723debf9383c75947
title: 4단계
challengeType: 0
dashedName: step-4
---

# --description--

`loop` 속성은 재생이 완료되면 비디오를 다시 시작합니다. 반복 재생되는 인터넷 밈을 생각해 보세요. `loop` 속성을 생략하면 비디오가 한 번만 재생됩니다.

`loop` 속성은 불리언 속성이며 값을 지정할 필요가 없습니다.

`loop` 속성을 `video` 요소에 추가하세요.

# --hints--

`video` 요소에 `loop` 속성이 있어야 합니다.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('loop'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640">
--fcc-editable-region--
  </video>
</body>
</html>
```
