---
id: 68f18724d71db1dda8d6fec7
title: 1단계
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

이 워크숍에서는 HTML 비디오 플레이어를 만듭니다. HTML 보일러플레이트가 제공되어 있습니다.

`h1` 요소를 만들고 텍스트로 `Working with the HTML Video Element`을 넣으세요.

# --hints--

`h1` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('h1'));
```

`h1` 요소 안에 `Working with the HTML Video Element` 텍스트가 있어야 합니다.

```js
const h1 = document.querySelector('h1');
assert.strictEqual(h1?.textContent.trim(), 'Working with the HTML Video Element');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
