---
id: 68f1eb9ad0a7575e4dbc31d5
title: 2단계
challengeType: 0
dashedName: step-2
---

# --description--

다음으로, `video` 아래에 `h1` 요소를 만드세요. 앞으로 몇 단계에 걸쳐 비디오 플레이어를 작동시키는 데 필요한 속성을 추가할 것입니다.

# --hints--

`video` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('video'));
```

새로운 `video` 요소는 `h1` 아래에 있어야 합니다.

```js
assert.exists(document.querySelector('h1 + video'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  
--fcc-editable-region--
</body>
</html>
```
