---
id: 68f1f8c248306da204655117
title: 6단계
challengeType: 0
dashedName: step-6
---

# --description--

`muted` 속성은 처음 재생할 때 오디오를 음소거합니다. `controls`이 활성화되어 있으면 사용자가 오디오 음소거를 해제할 수 있습니다. `muted` 속성을 생략하면 처음 재생할 때 오디오가 재생됩니다.

`muted` 속성은 불리언 속성이며 값이 필요 없습니다.

`muted` 요소에 `video` 속성을 추가하세요.

# --hints--

`video` 요소에 `muted` 속성이 있어야 합니다.

```js
const video = document.querySelector('video');
assert.isTrue(video?.hasAttribute('muted'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
--fcc-editable-region--
  <video width="640" loop controls>
--fcc-editable-region--
  </video>
</body>
</html>
```
