---
id: 68f26a5f79e97c5032c10642
title: 10단계
challengeType: 0
dashedName: step-10
---

# --description--

`mp4` 파일 확장자를 가진 비디오 파일을 사용했고, 브라우저가 파일을 어떻게 읽어야 할지 알 수 있도록 알려야 합니다.

`type` 속성을 사용해 `video/mp4` MIME 타입을 지정할 것입니다.

MIME(다목적 인터넷 메일 확장)은 ASCII 텍스트 외에 오디오, 비디오, 이미지 같은 문서 형식을 설명하는 표준입니다.

MP4는 공식 명칭으로 MPEG-4 Part 14이며, 디지털 멀티미디어 컨테이너 형식입니다. 비디오와 오디오 저장에 널리 사용되지만, 자막이나 정지 이미지 같은 다른 데이터 유형도 포함할 수 있습니다. MP4 파일은 인터넷 스트리밍용으로 설계되었으며, 많은 기기와 플랫폼에서 호환 가능합니다.

이제 `type` 속성과 값 `video/mp4`를 추가하세요.

# --hints--

`source` 요소에는 `type` 속성이 있어야 합니다.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('type'));
```

`source` 요소에는 값이 `type`인 `video/mp4` 속성이 있어야 합니다.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('type'), 'video/mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
--fcc-editable-region--
      
--fcc-editable-region--
    >
  </video>
</body>
</html>
```
