---
id: 68f268880219c2485bd71d2a
title: 9단계
challengeType: 0
dashedName: step-9
---

# --description--

비디오의 미디어 리소스를 지정하려면 `src` 요소에 `source` 속성을 추가해야 합니다.

값이 `src`인 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4` 속성을 추가하세요.

# --hints--

`source` 요소에 `src` 속성이 있어야 합니다.

```js
const source = document.querySelector('source');
assert.isTrue(source?.hasAttribute('src'));
```

`source` 요소에 값이 `src`인 `https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4` 속성이 있어야 합니다.

```js
const source = document.querySelector('source');
assert.strictEqual(source?.getAttribute('src'),
'https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
--fcc-editable-region--
    <source>
--fcc-editable-region--
  </video>
</body>
</html>
```
