---
id: 68f2701b2160c765f23b60f8
title: 12단계
challengeType: 0
dashedName: step-12
---

# --description--

또 다른 일반적인 MIME 유형은 `video/ogg` MIME 유형입니다.

Ogg는 디지털 멀티미디어를 효율적으로 스트리밍하고 조작할 수 있도록 설계된 디지털 멀티미디어 컨테이너 형식입니다. Xiph.Org 재단에서 관리하며 소프트웨어 특허에 제한받지 않는 무료 오픈 형식입니다. 이름은 컴퓨터 게임 Netrek의 전문 용어인 "ogging"에서 유래했습니다.

두 번째 `source` 요소 아래에 세 번째 `source` 요소를 추가하고 `src` 속성을 값 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg`로, `type` 속성을 값 `video/ogg`으로 지정하세요.

# --hints--

세 번째 `source` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('video > source:nth-of-type(3)'));
```

세 번째 `source` 요소에는 `src` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('src'));
```

세 번째 `source` 요소에는 값이 `src`인 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg');
```

세 번째 `source` 요소에는 `type` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.isTrue(source?.hasAttribute('type'));
```

세 번째 `source` 요소에는 값이 `type`인 `video/ogg` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(3)');
assert.strictEqual(source?.getAttribute('type'), 'video/ogg');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
