---
id: 68f26d52b13e0f5b6e2d2e09
title: 11단계
challengeType: 0
dashedName: step-11
---

# --description--

또 다른 일반적인 MIME 유형은 `video/webm` MIME 유형입니다.

WebM은 구글에서 개발한 오픈소스 시청각 미디어 파일 형식으로, 주로 웹 기반 미디어 콘텐츠를 위해 설계되었습니다. VP8, VP9, AV1 같은 비디오 코덱과 Vorbis, Opus 같은 오디오 코덱을 지원하여 HTML5 비디오 및 오디오 요소에 널리 사용됩니다.

첫 번째 `source` 요소 아래에 또 다른 `source` 요소를 추가하고 `src` 속성에 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` 값을, `type` 속성에 `video/webm` 값을 지정하세요.

# --hints--

두 번째 `source` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('video > source:nth-of-type(2)'));
```

두 번째 `source` 요소에는 `src` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('src'));
```

두 번째 `source` 요소에는 `src` 속성이 있고 그 값은 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm` 여야 합니다.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm');
```

두 번째 `source` 요소에는 `type` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.isTrue(source?.hasAttribute('type'));
```

두 번째 `source` 요소에는 `type` 속성이 있고 그 값은 `video/webm` 여야 합니다.

```js
const source = document.querySelector('source:nth-of-type(2)');
assert.strictEqual(source?.getAttribute('type'), 'video/webm');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```
