---
id: 68f27139bddf9f6bf2821acf
title: 13단계
challengeType: 0
dashedName: step-13
---

# --description--

마지막으로 추가할 `source` 요소는 `video/quicktime` MIME 유형에 해당합니다.

QuickTime은 Apple에서 만든 확장 가능한 멀티미디어 아키텍처로, 다양한 디지털 미디어 형식을 재생, 스트리밍, 인코딩, 트랜스코딩하는 것을 지원합니다. MP4 형식만큼 대중적이지는 않지만, 레거시 애플리케이션 지원을 위해 필요할 수 있습니다.

세 번째 `source` 요소 아래에 네 번째 `source` 요소를 추가하고 `src` 속성에 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov` 값을, `type` 속성에 `video/quicktime` 값을 지정하세요.

축하합니다! HTML 비디오 플레이어 워크숍을 완료했습니다.

# --hints--

네 번째 `source` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('video > source:nth-of-type(4)'));
```

네 번째 `source` 요소에는 `src` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('src'));
```

네 번째 `source` 요소에는 `src` 속성이 있고 값은 `https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov`이어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('src'), 'https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov');
```

네 번째 `source` 요소에는 `type` 속성이 있어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.isTrue(source?.hasAttribute('type'));
```

네 번째 `source` 요소에는 `type` 속성이 있고 값은 `video/quicktime`이어야 합니다.

```js
const source = document.querySelector('source:nth-of-type(4)');
assert.strictEqual(source?.getAttribute('type'), 'video/quicktime');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
  <meta charset="UTF-8">
  <meta name="viewport" content="width=device-width, initial-scale=1.0">
  <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    loop
    controls
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
--fcc-editable-region--
    
--fcc-editable-region--
  </video>
</body>
</html>
```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Working with the HTML Video Element</title>
</head>
<body>
  <h1>Working with the HTML Video Element</h1>
  <video
    width="640"
    controls
    loop
    muted
    poster="https://cdn.freecodecamp.org/curriculum/labs/past-event2.jpg"
  >
    <source
       src="https://cdn.freecodecamp.org/curriculum/labs/what-is-the-map-method-and-how-does-it-work.mp4"
      type="video/mp4"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.webm"
      type="video/webm"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.ogg"
      type="video/ogg"
    >
    <source
      src="https://cdn.freecodecamp.org/curriculum/labs/mapmethod.mov"
      type="video/quicktime"
    >
  </video>
</body>
</html>
```
