---
id: bad87fee1348bd9aedf08726
title: Utiliza código hexadecimal (hex code) para indicar colores específicos
challengeType: 0
videoUrl: 'https://scrimba.com/c/c8W9mHM'
forumTopicId: 18350
dashedName: use-hex-code-for-specific-colors
---

# --description--

¿Sabías que hay otras maneras de representar colores en CSS? Una de estas formas se llama código hexadecimal, o código hex (hex code) para abreviar.

Normalmente usamos <dfn>decimales</dfn>, o números en base 10, que usan los símbolos del 0 al 9 para cada dígito. Los <dfn>hexadecimales</dfn> (o <dfn>hex</dfn>) son números en base 16. Esto significa que usan dieciséis símbolos distintos. Al igual que los decimales, los símbolos del 0 al 9 representan los valores de cero a nueve. Luego A, B, C, D, E, F representan los valores de diez a quince. En conjunto, del 0 a la F pueden representar un dígito en hexadecimal, dándonos 16 valores posibles en total. Puedes encontrar más información sobre los <a href="https://www.freecodecamp.org/news/hexadecimal-number-system/" target="_blank" rel="noopener noreferrer nofollow">números hexadecimales aquí</a>.

En CSS, podemos usar 6 dígitos hexadecimales para representar colores, dos para cada uno de los componentes rojo (R), verde (G) y azul (B). Por ejemplo, `#000000` es negro y también es el valor más bajo posible. Puedes encontrar más información sobre el <a href="https://www.freecodecamp.org/news/rgb-color-html-and-css-guide/#whatisthergbcolormodel" target="_blank" rel="noopener noreferrer nofollow">sistema de color RGB aquí</a>.

```css
body {
  color: #000000;
}
```

# --instructions--

Reemplaza la palabra `black` en el "background-color" (color de fondo) de nuestro elemento `body` con su código hexadecimal equivalente, `#000000`.

# --hints--

Tu elemento `body` debe tener un `background-color` black (negro).

```js
const body = document.querySelector('body');
const backgroundColor = window.getComputedStyle(body)['background-color']; 
assert.strictEqual(backgroundColor, 'rgb(0, 0, 0)');
```

Debes usar el código hexadecimal correspondiente al color negro en lugar de la palabra reservada `black`.

```js
assert.match(code, /body\s*{(([\s\S]*;\s*?)|\s*?)background.*\s*:\s*?#000(000)?((\s*})|(;[\s\S]*?}))/gi);
```

# --seed--

## --seed-contents--

```html
<style>
  body {
    background-color: black;
  }
</style>
```

# --solutions--

```html
<style>
  body {
    background-color: #000000;
  }
</style>
```
