---
id: 66eaddd04a9e533fba689001
title: ブログ投稿カードのデザイン
challengeType: 25
dashedName: lab-blog-post-card
demoType: onClick
---

# --description--

このラボでは、背景やボーダーのスタイルを練習しながら、ブログ投稿カードを作成します。

**目的:** 以下のユーザーストーリーを満たし、すべてのテストに合格してラボを完了してください。

**ユーザーストーリー:**

1. `div` 要素に `blog-post-card` クラスを付けて、カードのすべての要素を保持してください。
2. `.blog-post-card` 要素の中に、有効な `alt` 属性とテキストを持つ画像要素があり、`post-img` クラスを付けてください。画像の `https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg` 属性には `src` を使えます。
3. `div` 要素の中に `post-content` があり、`.blog-post-card` クラスを付けて、以下を含めてください:
    - `h2` 要素に `post-title` クラスを付けて、ブログ投稿のタイトルのテキストを入れてください。
    - `p` 要素に `post-excerpt` クラスを付けて、ブログ投稿の要約テキストを入れてください。
    - `a` 要素に `read-more` クラスを付けて、テキスト `Read More` を入れてください。
4. `body` 要素の背景色を白以外の色に設定してください。
5. `.blog-post-card` 要素に以下のスタイルを適用してください:
    - 白い背景。
    - 角を丸くする。
    - 任意の幅。
    - 任意のテキスト整列。
6. `.post-img` 要素は、画像がカードの幅いっぱいに広がり、下にボーダーがあるようにスタイルを設定してください。
7. `.post-content` 要素は、カード内にパディングがあるようにスタイルを設定してください。
8. `.post-title` と `.post-excerpt` 要素は、デフォルト以外のテキスト色と全方向にマージンを持つようにしてください。
9. `.read-more` 要素はボタンのようにスタイルを設定し、以下を持つようにしてください:
    - デフォルト以外のテキスト色。
    - 背景色。
    - 全方向のマージン。
    - 全方向のパディング。
    - 角を丸くする。
    - `display` プロパティを `inline-block` に設定。
    - 背景色が変わるホバーエフェクト。

**注意:** スタイルシートをHTMLにリンクし、CSSを適用してください。

# --hints--

`div` に `blog-post-card` クラスを付けてください。

```js
assert.exists(document.querySelector("div.blog-post-card"));
```

`img` 要素に `class` が `post-img` であることを確認してください。画像にはテキストのある `alt` 属性と値のある `src` 属性が必要です。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const img = card.querySelector("img.post-img");
assert.exists(img);
assert.isString(img.alt);
assert.isNotEmpty(img.alt);
assert.isTrue(img.hasAttribute("src"));
assert.isNotEmpty(img.getAttribute("src"));
```

`div` に `post-content` クラスを付けてください。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);
const content = card.querySelector("div.post-content");
assert.exists(content);
```

`h2` 要素に `post-title` クラスを付けてください。ブログ投稿のタイトルのテキストがあることを確認してください。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postTitle = content.querySelector("h2.post-title");
assert.exists(postTitle);
assert.isString(postTitle.textContent);
assert.isNotEmpty(postTitle.textContent);
```

`p` に `post-excerpt` クラスを付けてください。ブログ投稿の要約テキストがあることを確認してください。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const postExcerpt = content.querySelector("p.post-excerpt");
assert.exists(postExcerpt);
```

`a` 要素に `class` が `read-more` であることを確認してください。

```js
const card = document.querySelector("div.blog-post-card");
assert.exists(card);

const content = card.querySelector("div.post-content");
assert.exists(content);

const readMore = content.querySelector("a.read-more");
assert.exists(readMore);
```

`.read-more` 要素に `Read More` のテキストがあることを確認してください。

```js
const el = document.querySelector("div.post-content a.read-more");
assert.exists(el);
assert.strictEqual(el.textContent.trim(), "Read More");
```

`.blog-post-card` 要素に `border-radius` プロパティが設定されていることを確認してください（`0` や負の値ではない）。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const borderRadius = getComputedStyle(card).borderRadius;
assert.isAbove(parseInt(borderRadius), 0);
```

`.blog-post-card` 要素は白い背景であることを確認してください。

```js
const card = document.querySelector('.blog-post-card');
assert.exists(card);

const backgroundColor = getComputedStyle(card).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.oneOf(backgroundColor.toLowerCase(), whiteColors);
```

`body` 要素は白以外の背景色であることを確認してください。

```js
const bodyBg = getComputedStyle(document.body).backgroundColor;
const whiteColors = ['rgb(255, 255, 255)', 'rgba(255, 255, 255, 1)', '#ffffff', '#fff', 'white'];

assert.isFalse(whiteColors.includes(bodyBg.toLowerCase()));
```

`.blog-post-card` をターゲットにして `width` プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("width"));
```

`.blog-post-card` をターゲットにして `text-align` プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".blog-post-card")?.getPropertyValue("text-align"));
```

`.post-content` をターゲットにして `padding` プロパティを設定してください。

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle(".post-content")?.getPropertyValue("padding"));
```

`.read-more` 要素にホバーエフェクトがあることを確認してください。

```js
const link = document.querySelector('.read-more');
assert.exists(link);

const d = new __helpers.CSSHelp(document).getStyle('.read-more:hover');
assert.isNotEmpty(d.backgroundColor);
```

`.read-more` をターゲットにして `color` プロパティを設定してください。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreTextColor = getComputedStyle(readMore).color;
assert.notStrictEqual(readMoreTextColor, 'rgba(0, 0, 0, 0)');
```

`.read-more` をターゲットにして `background-color` プロパティを設定してください。

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBackgroundColor = getComputedStyle(readMore).backgroundColor;
assert.notStrictEqual(readMoreBackgroundColor, 'rgba(0, 0, 0, 0)');
```

`.read-more` をターゲットにして `margin` プロパティを設定してください。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreMargin = getComputedStyle(readMore).marginTop;
assert.notStrictEqual(readMoreMargin, '0px');
```

`.read-more` をターゲットにして `display` プロパティを設定してください。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreDisplay = getComputedStyle(readMore).display;
assert.strictEqual(readMoreDisplay, 'inline-block');
```

`.read-more` をターゲットにして `border-radius` プロパティを設定してください。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMoreBorderRadius = getComputedStyle(readMore).borderRadius;
assert.notStrictEqual(readMoreBorderRadius, '0px');
```

`.read-more` をターゲットにして `padding` プロパティを設定してください。 

```js
const readMore = document.querySelector('.read-more');
assert.exists(readMore);

const readMorePadding = getComputedStyle(readMore).padding;
assert.notStrictEqual(readMorePadding, '0px');
```

`.post-img` 要素はカードの `width` を埋めるようにしてください。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgWidth = getComputedStyle(postImg).width;
const cardWidth = getComputedStyle(postImg.parentElement).width;
assert.closeTo(Number(imgWidth.replace(/px/ig, "")), Number(cardWidth.replace(/px/ig, "")), 15);
```

`.post-img` 要素は `border-bottom` の値を持つことを確認してください。

```js
const postImg = document.querySelector('.post-img');
assert.exists(postImg);

const imgBorderBottom = getComputedStyle(postImg).borderBottomWidth;
assert.notStrictEqual(imgBorderBottom, '0px');
```

`.post-title` と `.post-excerpt` 要素はマージンとデフォルト以外のテキスト色を持つことを確認してください。

```js
const title = document.querySelector('.post-title');
const excerpt = document.querySelector('.post-excerpt');

function isColorApplied(element) {
  const color = getComputedStyle(element).getPropertyValue('color');
  return color && color !== 'rgba(0, 0, 0, 0)' && color !== 'transparent'; 
}

function isMarginApplied(element) {
  const margin = getComputedStyle(element).getPropertyValue('margin');
  return margin && margin !== '0px'; 
}

assert.isTrue(isColorApplied(title));
assert.isTrue(isMarginApplied(title));

assert.isTrue(isColorApplied(excerpt));
assert.isTrue(isMarginApplied(excerpt));

const styles = new __helpers.CSSHelp(document);
// Are CSS selectors used to style the elements:
const postTitle = [".post-title", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
const postExcerpt = [".post-excerpt", ".post-excerpt, .post-title", ".post-title, .post-excerpt"].some(selector => {
  return !!styles.getStyle(selector)?.getPropVal('color');
});
assert.isTrue(postTitle);
assert.isTrue(postExcerpt);
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
</head>
<body>

</body>
</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">
<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Blog Post Card</title>
    <link rel="stylesheet" href="styles.css">
</head>
<body>
    <div class="blog-post-card">
        <img src="https://cdn.freecodecamp.org/curriculum/labs/cover-photo.jpg" alt="Blog Post Image" class="post-img">
        <div class="post-content">
            <h2 class="post-title">Learn Web Development in 2024</h2>
            <p class="post-excerpt">Stay ahead of the curve with the latest trends in web development. Discover what's new and exciting in 2024</p>
            <a href="#" class="read-more">Read More</a>
        </div>
    </div>
</body>
</html>
```

```css
body {
    display: flex;
    justify-content: center;
    align-items: center;
    height: 100vh;
    margin: 0;
    background-color: #f0f0f0;
    font-family: Arial, sans-serif;
}

.blog-post-card {
    background-color: white;
    border-radius: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    overflow: hidden;
    width: 350px;
    margin: 20px;
    text-align: center;
}

.post-img {
    width: 100%;
    height: auto;
    border-bottom: 5px solid #333;
}

.post-content {
    padding: 20px;
}

.post-title {
    color: #333;
    margin: 0 0 10px;
}

.post-excerpt {
    color: #667;
    margin: 10px 0;
}

.read-more {
    display: inline-block;
    margin-top: 10px;
    padding: 10px 15px;
    background-color: #333;
    color: white;
    text-decoration: none;
    border-radius: 5px;
}

.read-more:hover {
    background-color: #555;
}
```
