---
id: 686ccc2c8b967e17ab18d593
title: تنفيذ طريقة التقسيم الثنائي
challengeType: 27
dashedName: implement-the-bisection-method
---

# --description--

تستخدم طريقة التقسيم الثنائي، المعروفة أيضًا بطريقة البحث الثنائي، البحث الثنائي لإيجاد جذور دالة ذات قيم حقيقية. تعمل هذه الطريقة عن طريق تضييق فترة تحتوي على الجذر التربيعي حتى تصل إلى قيمة ضمن هامش خطأ محدد.

على سبيل المثال، إذا كان هامش الخطأ `0.01`، ستستمر طريقة التقسيم الثنائي في تقليل الفترة إلى النصف حتى يصبح الفرق بين الحد الأعلى والحد الأدنى أقل من أو يساوي `0.01`.

في هذا المختبر، ستنفذ دالة تستخدم طريقة التقسيم الثنائي لإيجاد الجذر التربيعي لعدد.

**الهدف:** أتمم قصص المستخدم أدناه واجعل جميع الاختبارات تنجح لإكمال المختبر.

**قصص المستخدم:**

1. يجب أن تعرف دالة باسم `square_root_bisection` بثلاث معلمات:
    - العدد الذي تريد إيجاد الجذر التربيعي له.
    - هامش الخطأ المقبول للنتيجة. يجب أن تحدد قيمة افتراضية لهامش الخطأ.
    - الحد الأقصى لعدد التكرارات التي ستُنفذ. يجب أن تحدد عددًا افتراضيًا للتكرارات.

1. يجب أن تقوم دالة `square_root_bisection` بما يلي:
    - ترفع استثناء `ValueError` مع الرسالة `Square root of negative number is not defined in real numbers` إذا كان العدد الممرر للدالة سالبًا.
    - للأعداد `0` و `1`، تطبع الرسالة: `The square root of [number] is [number]` وتُرجع العدد نفسه كجذر تربيعي.
    - لأي عدد موجب آخر، تطبع الجذر التربيعي التقريبي مع الرسالة: `The square root of [square_target] is approximately [root]` وتُرجع القيمة المحسوبة للجذر.
    - إذا لم يتحقق شرط هامش الخطأ لأي قيمة، تطبع رسالة فشل: `Failed to converge within [maximum] iterations` وتُرجع `None`.

**ملاحظة**: لا يمكنك استيراد أي وحدة في هذا المختبر.

# --hints--

يجب ألا تستورد أي وحدة.

```js
({ test: () => assert(runPython(`len(_Node(_code).find_imports()) == 0`)) })
```

يجب أن يكون لديك دالة باسم `square_root_bisection`.

```js
({ test: () => assert(runPython(`_Node(_code).has_function("square_root_bisection")`)) })
```

يجب أن تحتوي دالة `square_root_bisection` على ثلاث معلمات.

```js
({ test: () => runPython(`
    import inspect 
    sig = inspect.signature(square_root_bisection)
    assert len(sig.parameters) == 3
`) })
```

يجب أن تحدد قيمة افتراضية لهامش الخطأ والحد الأقصى لعدد التكرارات.

```js
({ test: () => runPython(`
try:
  import inspect 
  sig = inspect.signature(square_root_bisection)
  assert len(sig.parameters) == 3
  square_root_bisection(4)
except TypeError:
  assert False
`) })
```

يجب أن ترفع دالة `square_root_bisection` استثناء `ValueError` مع الرسالة `Square root of negative number is not defined in real numbers` عندما يكون العدد الممرر للدالة سالبًا.

```js
({ test: () => runPython(`
try:
  square_root_bisection(-6)
except ValueError as e:
  assert str(e) == "Square root of negative number is not defined in real numbers"
else:
  assert False
`) })
```

`square_root_bisection(0)` يجب أن تُرجع `0`.

```js
({ test: () => runPython(`assert square_root_bisection(0) == 0`) })
```

`square_root_bisection(0)` يجب أن تطبع `The square root of 0 is 0`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(0)
assert "The square root of 0 is 0" in _out
`) })
```

`square_root_bisection(0.001, 1e-7, 50)` يجب أن تُرجع عددًا بين `0.03162267660168379` و `0.031622876601683794`.

```js
({ test: () => runPython(`assert 0.03162267660168379 <= square_root_bisection(0.001, 1e-7, 50) <= 0.031622876601683794`) })
```

`square_root_bisection(0.001, 1e-7, 50)` يجب أن تطبع `The square root of 0.001 is approximately X`، حيث `X` عدد بين `0.03162267660168379` و `0.031622876601683794`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.001, 1e-7, 50)

assert 0.03162267660168379 <= _root <= 0.031622876601683794
assert f"The square root of 0.001 is approximately {_root}" in _out
`) })
```

`square_root_bisection(0.25, 1e-7, 50)` يجب أن تُرجع عددًا بين `0.4999999` و `0.5000001`.

```js
({ test: () => runPython(`assert 0.4999999 <= square_root_bisection(0.25, 1e-7, 50) <= 0.5000001`) })
```

`square_root_bisection(0.25, 1e-7, 50)` يجب أن تطبع `The square root of 0.25 is approximately X`، حيث `X` عدد بين `0.4999999` و `0.5000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(0.25, 1e-7, 50)

assert 0.4999999 <= _root <= 0.5000001
assert f"The square root of 0.25 is approximately {_root}" in _out
`) })
```

`square_root_bisection(1)` يجب أن تُرجع `1`.

```js
({ test: () => runPython(`assert square_root_bisection(1) == 1`) })
```

`square_root_bisection(1)` يجب أن تطبع `The square root of 1 is 1`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(1)
assert "The square root of 1 is 1" in _out
`) })
```

`square_root_bisection(81, 1e-3, 50)` يجب أن تُرجع عددًا بين `8.999` و `9.001`.

```js
({ test: () => runPython(`assert 8.999 <= square_root_bisection(81, 1e-3, 50) <= 9.001`) })
```

`square_root_bisection(81, 1e-3, 50)` يجب أن تطبع `The square root of 81 is approximately X`، حيث `X` عدد بين `8.999` و `9.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(81, 1e-3, 50)

assert 8.999 <= _root <= 9.001
assert f"The square root of 81 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-3, 100)` يجب أن تُرجع عددًا بين `14.999` و `15.001`.

```js
({ test: () => runPython(`assert 14.999 <= square_root_bisection(225, 1e-3, 100) <= 15.001`) })
```

`square_root_bisection(225, 1e-3, 100)` يجب أن تطبع `The square root of 225 is approximately X`، حيث `X` عدد بين `14.999` و `15.001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-3, 100)

assert 14.999 <= _root <= 15.001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-5, 100)` يجب أن تُرجع عددًا بين `14.99999` و `15.00001`.

```js
({ test: () => runPython(`assert 14.99999 <= square_root_bisection(225, 1e-5, 100) <= 15.00001`) })
```

`square_root_bisection(225, 1e-5, 100)` يجب أن تطبع `The square root of 225 is approximately X`، حيث `X` عدد بين `14.99999` و `15.00001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-5, 100)

assert 14.99999 <= _root <= 15.00001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 100)` يجب أن تُرجع عددًا بين `14.9999999` و `15.0000001`.

```js
({ test: () => runPython(`assert 14.9999999 <= square_root_bisection(225, 1e-7, 100) <= 15.0000001`) })
```

`square_root_bisection(225, 1e-7, 100)` يجب أن تطبع `The square root of 225 is approximately X`، حيث `X` عدد بين `14.9999999` و `15.0000001`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print

_root = square_root_bisection(225, 1e-7, 100)

assert 14.9999999 <= _root <= 15.0000001
assert f"The square root of 225 is approximately {_root}" in _out
`) })
```

`square_root_bisection(225, 1e-7, 10)` يجب أن تُرجع `None`.

```js
({ test: () => runPython(`assert square_root_bisection(225, 1e-7, 10) is None`) })
```

`square_root_bisection(225, 1e-7, 10)` يجب أن تطبع `Failed to converge within 10 iterations`.

```js
({ test: () => runPython(`
built_in_print = print
_out = []

def custom_print(*args, **kwargs):
  call_args = [arg for arg in args]
  _out.extend(call_args)

print = custom_print
square_root_bisection(225, 1e-7, 10)
assert "Failed to converge within 10 iterations" in _out
`) })
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def square_root_bisection(square_target, tolerance=1e-7, max_iterations=100):
    if square_target < 0:
        raise ValueError('Square root of negative number is not defined in real numbers')
    if square_target == 1:
        root = 1
        print(f'The square root of {square_target} is 1')
    elif square_target == 0:
        root = 0
        print(f'The square root of {square_target} is 0')

    else:
        low = square_target if square_target < 1 else 1
        high = 1 if square_target < 1 else square_target
        root = None

        for _ in range(max_iterations):
            mid = (low + high) / 2
            square_mid = mid**2

            if high - low <= tolerance:
                root = mid
                break

            elif square_mid < square_target:
                low = mid
            else:
                high = mid

        if root is None:
            print(f"Failed to converge within {max_iterations} iterations")

        else:   
            print(f'The square root of {square_target} is approximately {root}')

    return root
```
