---
id: 6909ed80e1c195f49564edac
title: الخطوة 1
challengeType: 1
dashedName: step-1
---

# --description--

في هذه الورشة، ستتعلم كيفية العمل مع عبارات الشرط وعوامل المقارنة.

ابدأ بتعريف متغير باسم `truthyOrFalsy` وقم بتعيين القيمة المنطقية `true` له.

# --hints--

يجب أن يكون لديك متغير باسم `truthyOrFalsy`.

```js
assert.exists(truthyOrFalsy);
```

يجب تعيين القيمة `true` إلى متغير `truthyOrFalsy`.

```js
assert.isTrue(truthyOrFalsy);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
