---
id: 690a562fd611e8155599ce98
title: الخطوة 3
challengeType: 1
dashedName: step-3
---

# --description--

غيّر تعيين `truthyOrFalsy` بتعيين قيمته إلى `"freeCodeCamp"`. ستظل ترى `true` مسجلة في وحدة التحكم لأن السلسلة النصية غير الفارغة هي قيمة صحيحة.

# --hints--

يجب أن تكون قيمة المتغير `truthyOrFalsy` هي `"freeCodeCamp"`.

```js
assert.strictEqual(truthyOrFalsy, "freeCodeCamp");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = true;
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
