---
id: 690a3ce03013be44d116136e
title: الخطوة 12
challengeType: 1
dashedName: step-12
---

# --description--

يمكن أن يحتوي بيان شرطي على بند `else`، الذي ينفذ شفرة عندما يكون شرط `if` خاطئًا. إليك مثال على بيان `if...else`:

```js
if (condition) {
  console.log("condition is truthy");
} else {
  console.log("condition is falsy");
}
```

أضف بند `else` إلى بيان `if` الموجود. داخل جسم بند `else`، سجّل في وحدة التحكم العبارة `"Timmy is not old enough to drive."`.

# --hints--

يجب أن يحتوي بيان `if` الثالث على بند `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s*else\s*\{/s);
```

يجب أن تسجل في وحدة التحكم العبارة `"Timmy is not old enough to drive."` داخل جسم بند `else`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}\s*else\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is not old enough to drive.\2\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

if (timmyAge >= 16) {
  console.log("Timmy is old enough to drive.");
--fcc-editable-region--
}
--fcc-editable-region--
```
