---
id: 690a110a76d0496d9e7c28eb
title: الخطوة 10
challengeType: 1
dashedName: step-10
---

# --description--

أعلن عن متغير آخر باسم `timmyAge` وقم بتعيين قيمته إلى `18`.

# --hints--

يجب أن يكون لديك متغير باسم `timmyAge`.

```js
assert.exists(timmyAge);
```

يجب تعيين القيمة `18` إلى متغير `timmyAge`.

```js
assert.strictEqual(timmyAge, 18);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

--fcc-editable-region--

--fcc-editable-region--
```
