---
id: 690a57095c4afb27bbe751b3
title: الخطوة 4
challengeType: 1
dashedName: step-4
---

# --description--

القيم الخاطئة هي القيم التي تُقيّم إلى `false` في سياق القيمة المنطقية. أمثلة على القيم الخاطئة هي:

- `""` (سلسلة نصية فارغة)
- `0`
- `false`
- `null`
- `undefined`
- `NaN`

ستتعرف على كل هذه القيم أثناء تقدمك في هذا المقرر.

في الوقت الحالي، غيّر تعيين `truthyOrFalsy` بجعل قيمته سلسلة نصية فارغة. بعد ذلك، سترى `false` مسجلة في وحدة التحكم.

# --hints--

يجب تعيين سلسلة نصية فارغة إلى `truthyOrFalsy`.

```js
assert.strictEqual(truthyOrFalsy, "");
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "freeCodeCamp";
--fcc-editable-region--

console.log(Boolean(truthyOrFalsy));
```
