---
id: 690a58cc4f828a49eab1b53a
title: الخطوة 6
challengeType: 1
dashedName: step-6
---

# --description--

أعلن عن متغير باسم `hasDeveloperJob` وقم بتعيين القيمة المنطقية `true` له.

# --hints--

يجب أن يكون لديك متغير باسم `hasDeveloperJob`.

```js
assert.exists(hasDeveloperJob);
```

يجب تعيين القيمة `true` للمتغير `hasDeveloperJob`.

```js
assert.isTrue(hasDeveloperJob);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
