---
id: 690a0df27ee9ee36d935d530
title: الخطوة 8
challengeType: 1
dashedName: step-8
---

# --description--

أعلن عن متغير آخر باسم `isTimmyAGamer` وقم بتعيين القيمة المنطقية `false` له.

# --hints--

يجب إعلان متغير باسم `isTimmyAGamer`.

```js
assert.exists(isTimmyAGamer);
```

يجب تعيين القيمة المنطقية `false` للمتغير `isTimmyAGamer`.

```js
assert.isFalse(isTimmyAGamer);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}
--fcc-editable-region--

--fcc-editable-region--
```
