---
id: 690a06833d3e83b5226ade78
title: الخطوة 7
challengeType: 1
dashedName: step-7
---

# --description--

تُستخدم العبارة الشرطية لتشغيل شفرة بناءً على شرط معين. تتكون من الكلمة المفتاحية `if`، تليها شرط محاط بأقواس عادية، وجسم محاط بأقواس متعرجة. إليك مثالًا على عبارة `if`:

```js
if (condition) {
  console.log("condition is truthy");
}
```

عندما تكون قيمة `condition` قيمة صحيحة (سلسلة نصية غير فارغة، رقم غير صفري، `true`، إلخ) تُنفذ الشفرة داخل جسم عبارة `if`.

أسفل متغيرك، أنشئ عبارة `if` تستخدم `hasDeveloperJob` كشرط لها. داخل جسم عبارة `if`، سجّل في وحدة التحكم النص `"Timmy is employed as a developer."`.

# --hints--

يجب أن يكون لديك عبارة `if`.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(.+\)\s*\{?/);
```

يجب أن تستخدم عبارة `if` الخاصة بك `hasDeveloperJob` كشرط لها.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{?/);
```

يجب أن تسجل في وحدة التحكم النص `"Timmy is employed as a developer."` داخل عبارة `if` الخاصة بك.

```js
assert.match(__helpers.removeJSComments(code), /\bif\s*\(\s*hasDeveloperJob\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is employed as a developer.\1\s*\)\s*;?\s*\}/);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;
--fcc-editable-region--

--fcc-editable-region--
```
