---
id: 690a57fef66aac39dffc61a2
title: الخطوة 5
challengeType: 1
dashedName: step-5
---

# --description--

الآن بعد أن عرفت ما هي القيم الصحيحة والقيم الخاطئة، احذف كل من إعلان المتغير وعبارة `console.log` من شفرتك.

# --hints--

يجب ألا يكون لديك متغير `truthyOrFalsy`.

```js
assert.notMatch(__helpers.removeJSComments(code), /const\s+truthyOrFalsy\s*=/);
```

يجب ألا يكون لديك `console.log(Boolean(truthyOrFalsy))` في شفرتك.

```js
assert.notMatch(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--
const truthyOrFalsy = "";

console.log(Boolean(truthyOrFalsy));
--fcc-editable-region--
```
