---
id: 690a51806fb743bc896c3af7
title: الخطوة 2
challengeType: 1
dashedName: step-2
---

# --description--

كما تعلمت في درس سابق، القيم المنطقية هي نوع بيانات يحتوي فقط على القيم `true` و `false`.

القيمة الصحيحة هي قيمة تُقيّم إلى `true` في سياق القيمة المنطقية. أمثلة على القيم الصحيحة هي:

- سلاسل نصية غير فارغة
- أرقام غير صفرية
- القيمة المنطقية `true`

وقيم أخرى ستتعرف عليها لاحقًا.

يمكنك استخدام الدالة `Boolean()` للتحقق من صحة القيمة. على سبيل المثال، `Boolean("Hello World!")` سترجع `true` لأن `"Hello World!"` قيمة صحيحة.

الآن، استدعِ `Boolean(truthyOrFalsy)` وسجل النتيجة في وحدة التحكم.

# --hints--

يجب أن تسجل `Boolean(truthyOrFalsy)` في وحدة التحكم.

```js
assert.match(__helpers.removeJSComments(code), /console\s*\.\s*log\s*\(\s*Boolean\s*\(\s*truthyOrFalsy\s*\)\s*\)\s*;?/);
```

# --seed--

## --seed-contents--

```js
const truthyOrFalsy = true;
--fcc-editable-region--

--fcc-editable-region--
```
