---
id: 690a3ad44874a928a2fe3077
title: الخطوة 11
challengeType: 1
dashedName: step-11
---

# --description--

تسمح لك عوامل المقارنة بمقارنة القيم وكتابة عبارات شرطية أكثر تعقيدًا.

```js
if (today == birthday) {
  console.log("Happy Birthday!");
}
```

على سبيل المثال، الشفرة أعلاه تسجل `"Happy Birthday!"` في وحدة التحكم فقط عندما تكون قيم المتغيرين `today` و `birthday` متساوية.

أنشئ جملة شرطية ثالثة باستخدام `if`. في شرطها، استخدم العامل `>=` للتحقق مما إذا كانت قيمة `timmyAge` أكبر من أو تساوي `16`.

داخل جسم جملة `if` الجديدة، سجل `"Timmy is old enough to drive."` في وحدة التحكم.

# --hints--

يجب أن تحتوي شفرتك على جملة `if` ثالثة.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(.+\)\s*\{.+\}/s);
```

يجب أن تستخدم جملة `if` الثالثة الشرط `timmyAge >= 16`.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{.+\}/s);
```

يجب أن تسجل `"Timmy is old enough to drive."` في وحدة التحكم داخل جملة `if` الثالثة.

```js
assert.match(__helpers.removeJSComments(code), /const\s+timmyAge\s*\=\s*.+\bif\s*\(\s*timmyAge\s*>=\s*16\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy is old enough to drive.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

if (isTimmyAGamer) {
  console.log("Timmy loves to play World of Warcraft.");
}

const timmyAge = 18;

--fcc-editable-region--

--fcc-editable-region--
```
