---
id: 690a0f7a5c22c15222a6aa5c
title: الخطوة 9
challengeType: 1
dashedName: step-9
---

# --description--

أنشئ الآن جملة `if` أخرى تستخدم `isTimmyAGamer` كشرط لها. داخل جسم جملة `if`، سجّل في وحدة التحكم العبارة `"Timmy loves to play World of Warcraft."`.

لاحظ أنك لن ترى أي شيء مسجلاً في وحدة التحكم لأن شرط جملة `if` الجديدة خاطئ.

# --hints--

يجب أن يكون لديك جملتا `if` في شفرتك.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(.+\)\s*\{.+\}/s);
```

يجب أن تستخدم جملة `if` الثانية الشرط `isTimmyAGamer`.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{.+\}/s);
```

يجب أن تسجل العبارة `"Timmy loves to play World of Warcraft."` في وحدة التحكم داخل جملة `if` الثانية.

```js
assert.match(__helpers.removeJSComments(code), /(?:\bif\s*\(.+\)\s*\{.+\}.*)\bif\s*\(\s*isTimmyAGamer\s*\)\s*\{\s*console\s*\.\s*log\s*\(\s*('|"|`)Timmy loves to play World of Warcraft.\1\s*\)\s*;?\s*\}/s);
```

# --seed--

## --seed-contents--

```js
const hasDeveloperJob = true;

if (hasDeveloperJob) {
  console.log("Timmy is employed as a developer.");
}

const isTimmyAGamer = false;

--fcc-editable-region--

--fcc-editable-region--
```
