---
id: 616d4a84b756d9c4b8255093
title: Step 11
challengeType: 0
dashedName: step-11
---

# --description--

It's time to add some color to the marker. Remember that one way to add color to an element is to use a <dfn>color keyword</dfn> like `black`, `cyan`, or `yellow`.

Come promemoria, ecco come fare riferimento alla classe `freecodecamp`:

```css
.freecodecamp {

}
```

Crea una nuova regola CSS destinata alla classe `marker` e assegna alla proprietà `background-color` il valore `red`.

**Nota:** Non vedrai alcun cambiamento dopo aver aggiunto il CSS.

# --hints--

Dovresti creare un selettore di classe per selezionare la classe `marker`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker'));
```

La regola CSS `.marker` dovrebbe avere una proprietà `background-color` con il valore `red`.

```js
assert(new __helpers.CSSHelp(document).getStyle('.marker')?.backgroundColor === 'red');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="utf-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Colored Markers</title>
    <link rel="stylesheet" href="styles.css">
  </head>
  <body>
    <h1>CSS Color Markers</h1>
    <div class="container">
      <div class="marker">
      </div>
    </div>
  </body>
</html>
```

```css
h1 {
  text-align: center;
}

--fcc-editable-region--

--fcc-editable-region--

```
