---
id: 6735a7370e0ae93a4577c689
title: 신문 레이아웃 디자인
challengeType: 25
dashedName: design-a-newspaper-layout
demoType: onClick
---

# --description--

이 실습에서는 CSS Grid를 사용해 신문 레이아웃을 디자인합니다. 여기에는 그리드 행과 그리드 열 같은 개념이 포함됩니다.

**목표:** 아래 사용자 스토리를 모두 충족하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. 페이지에는 클래스가 `main`인 `newspaper-layout` 요소가 있어야 합니다.
1. `.newspaper-layout`에는 신문 이름을 표시하는 `header`를 포함하는 클래스 `title`인 `h1`가 있어야 합니다.
1. `.newspaper-layout`에는 주요 뉴스 기사를 위한 클래스 `article`인 `feature-article`가 포함되어야 합니다.
1. `.feature-article`에는 기사 제목을 위한 `h2` 요소와 기사 내용을 위한 `p` 요소가 포함되어야 합니다.
1. 세 개의 작은 기사를 위한 `article` 요소를 각각 클래스 `small-article1`, `small-article2`, `small-article3`로 추가해야 합니다.
1. 각 작은 기사에는 기사 제목을 위한 `h3` 요소와 기사 내용을 위한 `p` 요소가 포함되어야 합니다.
1. `.newspaper-layout`에는 추가 뉴스 섹션을 위한 클래스 `article`인 `secondary-article` 요소가 포함되어야 합니다.
1. `.secondary-article`에는 기사 제목을 위한 `h2` 요소와 기사 내용을 위한 `p` 요소가 포함되어야 합니다.
1. `.newspaper-layout`에는 신문 내용을 나타내는 이미지를 표시하는 클래스 `figure`인 `cover-image`가 포함되어야 합니다.
1. 클래스가 `title`, `feature-article`, `secondary-article`, `cover-image`, `small-article1`, `small-article2`, `small-article3`인 요소들은 모두 같은 클래스 이름으로 `grid-area` 속성이 설정되어야 합니다.
1. `.newspaper-layout`는 CSS Grid를 사용하며 `grid-template-areas` 속성으로 섹션 배치를 정의해야 합니다:
   - `.title`는 상단 행 전체를 차지해야 합니다.
   - `.feature-article`는 두 번째 행에서 두 열을 차지하고 `.cover-image`는 세 번째 열에 있어야 합니다.
   - `.secondary-article`는 세 번째 행에서 두 열을 차지하고 `.cover-image`는 세 번째 열에 있어야 합니다.
   - 세 개의 작은 기사는 네 번째 행에 나타나야 합니다.
1. `.newspaper-layout`는 공간을 세 개의 동일한 열로 나누는 `grid-template-columns` 속성을 가져야 합니다.
1. `.newspaper-layout`의 `grid-template-rows` 속성은 첫 번째 행에 `auto`로 설정하고 나머지 공간은 동일한 크기로 나누어야 합니다.
1. 그리드 항목 사이에 간격을 추가해야 합니다.
1. `.cover-image` 안의 이미지는 지정된 공간에 맞도록 `max-width`를 `100%`로 설정해야 합니다.

**참고:** HTML에 스타일시트를 연결하고 CSS를 적용하는 것을 잊지 마세요.

# --hints--

`main` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector('main'));
```

`main` 요소는 `class`인 `newspaper-layout`를 가져야 합니다.

```js
assert.exists(document.querySelector('main.newspaper-layout'));
```

`header` 요소 안에 `class`인 `title`를 가진 `.newspaper-layout` 요소가 있어야 합니다.

```js
const header = document.querySelector(".newspaper-layout > header");
assert.exists(header);
assert(header.classList.contains("title"));
```

`header` 요소는 `h1` 요소를 포함해야 합니다.

```js
const h1 = document.querySelector("header h1");
assert.exists(h1);
```

`.title` 요소는 텍스트를 포함해야 합니다.

```js
assert.isAbove(document.querySelector(".title").innerText.length, 0);
```

각 기사에는 `article` 요소가 있어야 하며, 각각 클래스 `feature-article`, `small-article1`, `small-article2`, `small-article3`, `secondary-article`를 가져야 합니다.

```js
const featureArticle = document.querySelector('article.feature-article');
const smallArticle1 = document.querySelector('article.small-article1');
const smallArticle2 = document.querySelector('article.small-article2');
const smallArticle3 = document.querySelector('article.small-article3');
const secondaryArticle = document.querySelector('article.secondary-article');

assert.exists(featureArticle);
assert.exists(smallArticle1);
assert.exists(smallArticle2);
assert.exists(smallArticle3);
assert.exists(secondaryArticle);
```

`.feature-article`에는 `h2` 요소와 이어서 `p` 요소가 포함되어야 합니다.

```js
const h2 = document.querySelector(".feature-article h2");
const p = document.querySelector(".feature-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`.small-article1`, `.small-article2`, `.small-article3`에는 각각 `h3` 요소와 이어서 `p` 요소가 포함되어야 합니다.

```js
for (let i = 1; i < 4; i++) {
    let h3 = document.querySelector(`.small-article${i} h3`);
    let p = document.querySelector(`.small-article${i} h3 + p`);
    assert.exists(h3);
    assert.exists(p);
}
```

`.secondary-article`에는 `h2` 요소와 이어서 `p` 요소가 포함되어야 합니다.

```js
const h2 = document.querySelector(".secondary-article h2");
const p = document.querySelector(".secondary-article h2 + p");
assert.exists(h2);
assert.exists(p);
```

`figure` 안에 클래스 `cover-image`인 `.newspaper-layout` 요소가 있어야 합니다.

```js
assert.exists(document.querySelector(".newspaper-layout figure.cover-image"));
```

`.cover-image`는 `img` 요소를 포함해야 합니다.

```js
assert.exists(document.querySelector(".cover-image img"));
```

`display` 요소의 `.newspaper-layout` 속성은 `grid`로 설정되어야 합니다.

```js
assert.equal(getComputedStyle(document.querySelector('.newspaper-layout')).display, 'grid');
```

`.newspaper-layout` 요소는 첫 번째 행의 세 열 전체를 차지하는 `grid-template-areas`가 있는 `title` 속성을 가져야 합니다.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
assert.match(newspaperLayout.gridTemplateAreas, /^("|')title title title\1/);
```

그리드 템플릿의 두 번째 행에는 두 열을 차지하는 `feature-article`와 `cover-image`가 포함되어야 합니다.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromSecondRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").trimStart();
assert.match(fromSecondRow, /^("|')feature-article feature-article cover-image\1/);
```

그리드 템플릿의 세 번째 행에는 두 열을 차지하는 `secondary-article`와 `cover-image`가 포함되어야 합니다.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromThirdRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromThirdRow, /^("|')secondary-article secondary-article cover-image\1/);
```

그리드 템플릿의 네 번째 행에는 `small-article1`, `small-article2`, `small-article3`가 포함되어야 합니다.

```js
const newspaperLayout = getComputedStyle(document.querySelector('.newspaper-layout'));
const rowRegex = /("|')(.+?)\1/;
const fromFourthRow = newspaperLayout.gridTemplateAreas.replace(rowRegex, "").replace(rowRegex, "").replace(rowRegex, "").trimStart();
assert.match(fromFourthRow, /^("|')small-article1 small-article2 small-article3\1$/);
```

`.newspaper-layout` 선택자는 공간을 세 개의 동일한 부분으로 나누는 `grid-template-columns` 속성을 사용해야 합니다. `repeat` 함수나 각 분수 단위를 수동으로 정의할 수 있습니다.

```js
assert.oneOf(
    new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateColumns,
    ['1fr 1fr 1fr', 'repeat(3, 1fr)'] 
);
```

`.newspaper-layout` 선택자는 `grid-template-rows`로 설정된 `auto 1fr 1fr 1fr` 속성을 가져야 합니다.

```js
assert.oneOf(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gridTemplateRows, ['auto 1fr 1fr 1fr', 'auto repeat(3, 1fr)']);
```

그리드 항목 사이에 간격을 추가해야 합니다.

```js
assert.isNotEmpty(new __helpers.CSSHelp(document).getStyle('.newspaper-layout')?.gap);
```

`.title`는 `grid-area`인 `title`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.title')?.gridArea, 'title');
```

`.feature-article`는 `grid-area`인 `feature-article`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.feature-article')?.gridArea, 'feature-article');
```

`.cover-image`는 `grid-area`인 `cover-image`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.cover-image')?.gridArea, 'cover-image');
```

`.secondary-article`는 `grid-area`인 `secondary-article`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.secondary-article')?.gridArea, 'secondary-article');
```

`.small-article1`는 `grid-area`인 `small-article1`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article1')?.gridArea, 'small-article1');
```

`.small-article2`는 `grid-area`인 `small-article2`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article2')?.gridArea, 'small-article2');
```

`.small-article3`는 `grid-area`인 `small-article3`를 가져야 합니다.

```js
assert.equal(new __helpers.CSSHelp(document).getStyle('.small-article3')?.gridArea, 'small-article3');
```

`img` 안의 `.cover-image`는 `max-width`인 `100%`를 가져야 합니다.

```js
const coverImage = document.querySelector(".cover-image");
const img = document.querySelector(".cover-image img");
assert.equal(getComputedStyle(img).maxWidth, "100%");
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">

<head>
<meta charset="utf-8">
<title>Newspaper Layout</title>
</head>

<body>

</body>

</html>
```

```css

```

# --solutions--

```html
<!DOCTYPE html>
<html lang="en">

<head>
    <meta charset="UTF-8">
    <meta name="viewport" content="width=device-width, initial-scale=1.0">
    <title>Newspaper Layout</title>
    <link rel="stylesheet" href="styles.css">
</head>

<body>
    <main class="newspaper-layout">
        <header class="title">
            <h1>The Daily Local News</h1>
        </header>

        <article class="feature-article">
            <h2>Breaking News: Volcano Eruption Disrupts Tourism</h2>
            <p>Recently, a volcano erupted in a popular tourist destination. The eruption has caused widespread panic
                and has disrupted tourism in the area. The volcano has been spewing lava and ash for several days now,
                and authorities are urging residents and tourists to evacuate the area immediately. The eruption has
                also caused a number of flights to be cancelled, leaving many tourists stranded. The situation is still
                developing, and authorities are working to contain the eruption and ensure the safety of everyone in the
                area.</p>
        </article>

        <article class="small-article1">
            <h3>Sports: Local Team Wins Championship</h3>
            <p>Hockey fans are celebrating today as the local team has won the championship. The team, which has been on
                a winning streak all season, clinched the title in a thrilling final match. Fans took to the streets to
                celebrate the victory, waving flags and chanting the team's name.</p>
        </article>

        <article class="small-article2">
            <h3>Health: Tips for a Balanced Diet</h3>
            <p>A diet high in calories, sugar, and unhealthy fats can lead to a variety of health problems, including
                obesity, diabetes, and heart disease. To maintain a healthy weight and reduce your risk of chronic
                diseases, it's important to eat a balanced diet that includes a variety of nutrient-rich foods. Here are
                some tips for eating a balanced diet:</p>
        </article>

        <article class="small-article3">
            <h3>Travel: Top 10 Destinations for 2025</h3>
            <p>Traveling is one of the best ways to experience new cultures, meet new people, and see the world. If
                you're looking for inspiration for your next trip, here are the top 10 destinations for 2025:</p>
        </article>

        <article class="secondary-article">
            <h2>Technology: The Rise of AI</h2>
            <p>Artificial intelligence (AI) is changing the way we live and work. From self-driving cars to virtual
                assistants, AI is revolutionizing the world around us. But what exactly is AI, and how does it work? In
                this article, we'll explore the rise of AI and its impact on society.</p>
        </article>

        <figure class="cover-image">
            <img src="https://cdn.freecodecamp.org/curriculum/labs/volcano.jpg" alt="Cover Image">
        </figure>
    </main>
</body>

</html>
```

```css
body {
    font-family: 'Georgia', 'Times New Roman', Times, serif;
    padding: 0;
    margin: 0;
    background-color: #f0f0f0;
}

.newspaper-layout {
    display: grid;
    grid-template-areas:
        "title title title"
        "feature-article feature-article cover-image"
        "secondary-article secondary-article cover-image"
        "small-article1 small-article2 small-article3";
    grid-template-columns: 1fr 1fr 1fr;
    grid-template-rows: auto 1fr 1fr 1fr;
    gap: 15px;
    background-color: #fff;
    padding: 15px;
    box-shadow: 0 4px 8px rgba(0, 0, 0, 0.1);
    box-sizing: border-box;
    text-align: justify;
}

.title {
    grid-area: title;
    background-color: #1B1B32;
    color: white;
    padding: 15px;
    text-align: center;
}

.title h1 {
    font-size: 2.5rem;
    line-height: 1.3;
    margin: 0;
}

.feature-article,
.secondary-article {
    background-color: #f9f9f9;
    padding: 15px;
    font-size: 1rem;
}

.feature-article {
    grid-area: feature-article;
}

.secondary-article {
    grid-area: secondary-article;
}

.small-article1,
.small-article2,
.small-article3 {
    background-color: #e9e9e9;
    padding: 15px;
    font-size: 0.9rem;
}

.small-article1 {
    grid-area: small-article1;
}

.small-article2 {
    grid-area: small-article2;
}

.small-article3 {
    grid-area: small-article3;
}

.cover-image {
    grid-area: cover-image;
    display: flex;
    justify-content: center;
    align-items: center;
    background-color: #ddd;
    padding: 15px;
    margin: 0;
}

.cover-image img {
    max-width: 100%;
    height: 350px;
    width: 350px;
    border-radius: 5px;
}
```
