---
id: 66f17db06803d11a1bd19a20
title: クイズゲームを作ろう
challengeType: 26
dashedName: lab-quiz-game
---

# --description--

**目的:** 以下のストーリーを満たし、すべてのテストに合格してラボを完成させてください。

**ユーザーストーリー:**

1. `questions` という名前の配列を作成してください。
1. `questions` 配列には、`category`、`question`、`choices`、`answer` というキーを持つオブジェクトが少なくとも5つ含まれている必要があります。
1. `category` キーには、質問のカテゴリを表す文字列の値を設定してください。
1. `question` キーには、質問を表す文字列の値を設定してください。
1. `choices` キーには、質問の選択肢となる3つの文字列を含む配列の値を設定してください。
1. `answer` キーには、質問の正しい答えを表す文字列の値を設定してください。また、`answer` の値は `choices` 配列に含まれている必要があります。
1. `getRandomQuestion` という名前の関数を作成してください。この関数は質問の配列をパラメータとして受け取り、その配列からランダムな質問オブジェクトを返します。
1. `getRandomComputerChoice` という名前の関数を作成してください。この関数は利用可能な選択肢の配列をパラメータとして受け取り、選択された質問のランダムな答えを返します。
1. `getResults` という名前の関数を作成してください。この関数は質問オブジェクトを第一パラメータに、コンピューターの選択を第二パラメータに受け取ります。答えが正しければ `The computer's choice is correct!` を返し、そうでなければ `The computer's choice is wrong. The correct answer is: <correct-answer>` を返します。ただし `<correct-answer>` は選択された質問の正しい答えの値です。

# --hints--

`questions` という名前の配列を作成してください。

```js
assert.isArray(questions);
```

`questions` 配列には、`category`、`question`、`choices`、`answer` というキーを持つオブジェクトが少なくとも5つ含まれている必要があります。

```js
assert.isAtLeast(questions.length, 5);
questions.forEach(question => {
    assert.isObject(question);
    assert.hasAllKeys(question, ['category', 'choices', 'question', 'answer']);
});

```

`category` キーには、質問のカテゴリを表す文字列の値を設定してください。

```js
assert.isNotEmpty(questions);
questions.forEach(({category}) => {
    assert.isString(category);
    assert.isNotEmpty(category);
});
```

`question` キーには、質問を表す文字列の値を設定してください。

```js
assert.isNotEmpty(questions);
questions.forEach(({question}) => {
    assert.isString(question);
    assert.include(question, '?');
    assert.isNotEmpty(question.replaceAll('?', ''));
})
```

`choices` キーには、互いに異なる3つの文字列を含む配列の値を設定してください。

```js
assert.isNotEmpty(questions);
questions.forEach(({choices}) => {
    assert.isArray(choices);
    assert.lengthOf(choices, 3);
    assert.lengthOf(new Set(choices), 3);
    choices.forEach(choice => {
        assert.isString(choice);
        assert.isNotEmpty(choice);
    });
})
```

`answer` キーには文字列の値を設定してください。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer}) => {
    assert.isString(answer);
    assert.isNotEmpty(answer);
});
```

`answer` の値は `choices` 配列に含まれている必要があります。

```js
assert.isNotEmpty(questions);
questions.forEach(({answer, choices}) => assert.oneOf(answer, choices));
```

`getRandomQuestion` という名前の関数を作成してください。この関数は質問の配列をパラメータとして受け取り、その配列からランダムな質問オブジェクトを返します。

```js
const parameters = __helpers.getFunctionParams(getRandomQuestion.toString());

assert.isFunction(getRandomQuestion);
assert.lengthOf(parameters, 1);
assert.lengthOf(getRandomQuestion, 1);
assert.isObject(getRandomQuestion(questions));
assert.deepInclude(questions, getRandomQuestion(questions), "getRandomQuestion did not return one of the objects inside questions");
```

`getRandomComputerChoice` という名前の関数を作成してください。この関数は利用可能な選択肢の配列をパラメータとして受け取り、選択された質問のランダムな答えを返します。

```js
assert.isFunction(getRandomComputerChoice);
const randomArray = ['a', 'b', 'c'];
assert.oneOf(getRandomComputerChoice(randomArray), randomArray)
```

`getResults` という名前の関数を作成してください。

```js
assert.isFunction(getResults);
```

`getResults` 関数は質問オブジェクトを第一パラメータに、コンピューターの選択を第二パラメータに受け取る必要があります。

```js
const testQuestion = questions[0];
const wrongChoice = questions[0].choices.find(choice => choice !== questions[0].answer);
assert.equal(getResults(testQuestion, testQuestion.answer), `The computer's choice is correct!`);
assert.equal(getResults(testQuestion, wrongChoice), `The computer's choice is wrong. The correct answer is: ${testQuestion.answer}`);
```

コンピューターの選択が答えと一致する場合、`getResults` は `The computer's choice is correct!` を返す必要があります。

```js
assert.equal(getResults(questions[0], questions[0].answer), `The computer's choice is correct!`);
```

コンピューターの選択が答えと一致しない場合、`getResults` は `The computer's choice is wrong. The correct answer is: <correct-answer>` を返す必要があります。ただし `<correct-answer>` は選択された質問の正しい答えの値です。

```js
assert.equal(getResults({category: 'misc', choices: ['a', 'b', 'c'], question: "question?", answer: "b"}, "a"), `The computer's choice is wrong. The correct answer is: b`)
```

`getResults` 関数は部分文字列の一致ではなく、厳密な等価比較を使う必要があります。

```js
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Hamburger"}, "Ham"), `The computer's choice is wrong. The correct answer is: Hamburger`);
assert.equal(getResults({category: 'food', choices: ['Ham', 'Hamburger', 'Hot Dog'], question: "What food?", answer: "Ham"}, "Hamburger"), `The computer's choice is wrong. The correct answer is: Ham`);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
const questions = [
    {
        category: "science",
        question: "What is the chemical symbol for potassium?",
        choices: ["P", "K", "Pt"],
        answer: "K"
    },
    {
        category: "science",
        question: "What is the unit of electrical resistance?",
        choices: ["Ohm", "Coulomb", "Sievert"],
        answer: "Ohm"
    },
    {
        category: "geography",
        question: "What is the capital city of Australia?",
        choices: ["Sidney", "Canberra", "Wellington"],
        answer: "Canberra"
    },
    {
        category: "literature",
        question: 'Who wrote "1984"?',
        choices: ["Ray Bradbury", "Aldous Huxley", "George Orwell"],
        answer: "George Orwell"
    },
    {
        category: "sport",
        question: "How many players are on a standard volleyball team?",
        choices: ["6", "7", "12"],
        answer: "6"
    }
]

function getRandomQuestion(questionList) {
    const randomIndex = Math.floor(Math.random() * questionList.length);
    return questionList[randomIndex];
}

function getRandomComputerChoice(choices) {
    const randomIndex = Math.floor(Math.random() * choices.length);
    return choices[randomIndex];
}

function getResults(obj, computerChoice) {
    return computerChoice === obj.answer
        ? "The computer's choice is correct!"
        : `The computer's choice is wrong. The correct answer is: ${obj.answer}`;
}

const questionObj = getRandomQuestion(questions);
const { question, choices } = questionObj;

console.log(question);
console.log(`Choices: ${choices}`);

const computerChoice = getRandomComputerChoice(choices);
console.log(`Computer chooses: ${computerChoice}`);

const results = getResults(questionObj, computerChoice);
console.log(results);

```
