---
id: 695774002591bbc5f8cf3e53
title: 割引適用関数を作成する
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

このラボでは、パーセント割引を適用した後の商品の最終価格を計算する関数を書きます。

たとえば、商品の価格が `50` で割引が `20` 適用された場合、割引額は `10` で、最終価格は `40` になります。

**目的:** 以下のユーザーストーリーを満たしてラボを完成させてください。

**ユーザーストーリー:**

1. `apply_discount` という名前の関数を定義してください。

2. `apply_discount` 関数はちょうど2つのパラメータ、`price` と `discount` を受け取る必要があります。

3. `price` が数値（`int` または `float`）でない場合、関数は文字列 `The price should be a number` を返す必要があります。

4. `discount` が数値（`int` または `float`）でない場合、関数は文字列 `The discount should be a number` を返す必要があります。

5. `price` が `0` 以下の場合、関数は文字列 `The price should be greater than 0` を返す必要があります。

6. `discount` が `0` 未満または `100` より大きい場合、関数は文字列 `The discount should be between 0 and 100` を返す必要があります。

7. 両方の入力が有効な場合、関数は価格のパーセントとして割引を計算する必要があります。

8. 関数は割引を適用した後の最終価格を返す必要があります。

# --hints--

`apply_discount` という名前の関数を用意してください。

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

`apply_discount` 関数は2つのパラメータ、`price` と `discount` を受け取る必要があります。

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

`apply_discount` が数値（`price` または `int`）でない `float`（最初の引数）で呼び出された場合、`The price should be a number` を返す必要があります。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

`apply_discount` が数値（`discount` または `int`）でない `float`（2番目の引数）で呼び出された場合、`The discount should be a number` を返す必要があります。

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

`apply_discount` が `price` 以下の `0` で呼び出された場合、`The price should be greater than 0` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

`apply_discount` が `discount` 未満または `0` より大きい `100` で呼び出された場合、`The discount should be between 0 and 100` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` は `80` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` は `100` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` は `50` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

`apply_discount` が `100` の割引で呼び出された場合、`0` を返す必要があります。

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` は `59.6` を返す必要があります。

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
