---
id: 6853e58d52ec64349a5760ca
title: Hatua 1
challengeType: 20
dashedName: step-1
---

# --description--

Katika warsha hii, utaekeleza algoriti ya njia fupi zaidi. Utaandika kitendakazi cha Python kinachohesabu njia fupi zaidi kati ya nodi katika grafu, na pia kurudisha njia iliyochukuliwa.

Kwa mfano, ukizingatia grafu ambapo miji imeunganishwa na barabara zenye umbali tofauti, algoriti itapata njia fupi zaidi kutoka mji mmoja hadi mwingine. Ikiwa unataka kusafiri kutoka Mji A hadi Mji D, algoriti inaweza kugundua kuwa kwenda A ⇨ B ⇨ C ⇨ D (jumla: km 15) ni fupi zaidi kuliko kwenda moja kwa moja A ⇨ D (km 20).

Ili kuanza, fafanua kigezo kinachoitwa `INF` na ukaweke thamani `float('inf')`, inayowakilisha ukomo chanya usio na kikomo. Baadaye, utaitumia kuonyesha umbali usio na kikomo kati ya nodi mbili.

# --hints--

Unapaswa kufafanua kigezo `INF` juu ya faili yako.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("INF")`));
  }
});
```

Unapaswa kuweka thamani `float('inf')` kwa kigezo `INF`.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("INF").is_equivalent("INF = float('inf')")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
