---
id: 686253f1b051998ad4904e3e
title: Hatua 3
challengeType: 20
dashedName: step-3
---

# --description--

Sasa utaunda kitendakazi kikuu kinachokubali vigezo vitatu: matriki ya ujirani, nodi ya kuanzia, na nodi lengwa ya hiari.

Unda kitendakazi kinachoitwa `shortest_path` kinachopokea vigezo vitatu: `matrix`, `start_node`, na `target_node`. Weka `None` kama thamani ya chaguo-msingi kwa `target_node`.

Kigezo cha `target_node` kimewekwa kuwa `None` kwa chaguo-msingi, kinachoashiria kwamba kama nodi lengwa haijaelezwa, kitendakazi kinapaswa kuhesabu njia fupi kutoka nodi ya kuanzia hadi nodi zote nyingine katika grafu.

Ongeza tamko la `pass` ndani ya mwili wa kitendakazi kwa sasa.

# --hints--

Unapaswa kuunda kitendakazi kinachoitwa `shortest_path`.

```js
({
    test: () => {
      assert(runPython(`_Node(_code).has_function("shortest_path")`));
    }
});
```

Kitendakazi kinapaswa kupokea vigezo vitatu: `matrix`, `start_node`, na `target_node` huku `target_node` ikiwa na thamani ya chaguo-msingi ya `None`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_args("matrix, start_node, target_node=None")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--

--fcc-editable-region--
```
