---
id: 68625ad34aa3f58d7f3b9b07
title: Hatua 4
challengeType: 20
dashedName: step-4
---

# --description--

Unahitaji kuhifadhi idadi ya nodi katika grafu. Kwa hili utahitaji kigezo kinacholingana na urefu wa matriki ya ujirani.

Ndani ya kitendakazi cha `shortest_path`, tengeneza kigezo `n` na uweke thamani yake kuwa urefu wa `matrix`.

# --hints--

Unapaswa kuanzisha `n` kwa urefu wa `matrix`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("n").is_equivalent("n = len(matrix)")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
--fcc-editable-region--
    pass
--fcc-editable-region--
```
