---
id: 686251fb49ebc489ba60ccb5
title: Hatua 2
challengeType: 20
dashedName: step-2
---

# --description--

Utahitaji kuunda orodha ya 2D ili wakilisha matriki ya ujirani ya grafu. Matriki hii itatumika kuwakilisha uzito wa makali kati ya nodi katika grafu.

Unda kigezo kinachoitwa `adj_matrix` na ukaweke thamani ya orodha ya 2D inayowakilisha grafu yenye uzito ufuatao:

```py
[0, 5, 3, INF, 11, INF],
[5, 0, 1, INF, INF, 2],
[3, 1, 0, 1, 5, INF],
[INF, INF, 1, 0, 9, 3],
[11, INF, 5, 9, 0, INF],
[INF, 2, INF, 3, INF, 0]
```

# --hints--

Unapaswa kuunda kigezo kinachoitwa `adj_matrix`.

```js
({
  test: () => {
    assert(runPython(`_Node(_code).has_variable("adj_matrix")`));
  }
});
```

`adj_matrix` inapaswa kuwa orodha ya 2D yenye thamani zilizotolewa.

```js
({
  test: () => {
    assert(
      runPython(
        `_Node(_code).find_variable("adj_matrix").is_equivalent("adj_matrix = [[0, 5, 3, INF, 11, INF], [5, 0, 1, INF, INF, 2], [3, 1, 0, 1, 5, INF], [INF, INF, 1, 0, 9, 3], [11, INF, 5, 9, 0, INF], [INF, 2, INF, 3, INF, 0]]")`
      )
    );
  }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
--fcc-editable-region--

--fcc-editable-region--
```
