---
id: 686b78bc13d518e1267448e4
title: Hatua 8
challengeType: 20
dashedName: step-8
---

# --description--

Mbali na kufuatilia umbali, pia unahitaji kufuatilia njia halisi zilizochukuliwa kufikia kila nodi.

Utatengeneza orodha ambapo kila kipengee kinahifadhi njia iliyochukuliwa kufikia nodi hiyo. Awali, njia ya kila nodi itakuwa na nodi yenyewe tu.

Orodha za kuelezea (list comprehensions) hutoa njia fupi ya kutengeneza orodha. Kwa mfano:

```py
[x * 2 for x in range(3)]
```

Tengeneza kigezo kinachoitwa `paths` na uanze kwa kutumia orodha ya kuelezea inayotengeneza orodha yenye `[node_no]` kwa kila `node_no` katika `range(n)`.

# --hints--

Unapaswa kuwa na kigezo kinachoitwa `paths`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("paths")`
        )
      );
    }
});
```

Unapaswa kuanzisha `paths` kwa kutumia orodha ya kuelezea.

```js
({  
    test: () => {  
      runPython(`  
import ast  
assert isinstance(_Node(_code).find_function("shortest_path").find_variable("paths").tree.value, ast.ListComp)  
    `);  
    }  
});
```

Orodha yako ya kuelezea inapaswa kutumia `node_no` kupita juu ya `range(n)`.

```js
({  
  test: () => {  
    runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert len(v.find_comp_targets()) == 1  
assert v.find_comp_targets()[0].is_equivalent("node_no")  
assert len(v.find_comp_iters()) == 1  
assert v.find_comp_iters()[0].is_equivalent("range(n)")  
    `);  
  }  
});
```

Orodha yako ya kuelezea inapaswa kutathmini `[node_no]` kwa kila `node_no` katika `range(n)`.

```js  
({  
    test: () => {  
      runPython(`  
v = _Node(_code).find_function("shortest_path").find_variable("paths")  
assert v.find_comp_expr().is_equivalent("[node_no]")  
    `);  
    }  
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
--fcc-editable-region--
    
--fcc-editable-region--
```
