---
id: 686b73007fc582db0132c361
title: Hatua 5
challengeType: 20
dashedName: step-5
---

# --description--

Unahitaji kufuatilia umbali mfupi zaidi unaojulikana kutoka nodi ya mwanzo hadi kila nodi nyingine katika grafu.

Ili kufanya hivyo, tengeneza kigezo kinachoitwa `distances` na uanze kwa kuiweka kama orodha yenye kipengele kimoja tu: `INF`.


# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `distances`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").has_variable("distances")`
        )
      );
    }
});
```

Orodha ya `distances` inapaswa kuanzishwa kuwa `[INF]`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF]")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
--fcc-editable-region--
    
--fcc-editable-region--
```
