---
id: 68777816e95d4936f129819f
title: Hatua 17
challengeType: 20
dashedName: step-17
---

# --description--

Kabla ya kujaribu kusahisha umbali kwa jirani, unahitaji kuthibitisha kuwa jirani huyo anapatikana na hajatembelewa. Kisha utahesabu jumla ya umbali itakavyokuwa kufikia jirani huyo kupitia nodi ya sasa.

Ndani ya mzunguko wa `for`, ongeza tamko la `if` linalokagua:

- Kama `distance` si sawa na `INF` (maana kuna kivuko kati ya nodi hizo)
- Jirani wa `node_no` hajatembelewa bado

Ndani ya blokki ya `if`, tengeneza kigezo kinachoitwa `new_distance` na ukaweke thamani ya jumla ya `distances[current]` (umbali mfupi zaidi hadi nodi ya sasa) na `distance` (umbali kutoka nodi ya sasa hadi jirani).

# --hints--

Unapaswa kuwa na tamko la `if` linalokagua kama `distance` si sawa na `INF` na kama nodi ya jirani hajatembelewa.

```js
({
    test: () => {
      assert(
        runPython(`_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_conditions()[0].is_equivalent("distance != INF and not visited[node_no]") 
      `)
      );
    }
});
```

Unapaswa kutengeneza kigezo kinachoitwa `new_distance` ndani ya blokki ya `if`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].has_variable("new_distance")`
        )
      );
    }
});
```

Unapaswa kuweka thamani ya `new_distance` kuwa jumla ya `distances[current]` na `distance`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_for_loops()[0].find_bodies()[0].find_for_loops()[1].find_bodies()[0].find_ifs()[0].find_bodies()[0].find_variable("new_distance").is_equivalent("new_distance = distances[current] + distance")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n

    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
              min_distance = distances[node_no]
              current = node_no

        if current == -1:
            break
        visited[current] = True

--fcc-editable-region--
        for node_no in range(n):  
            distance = matrix[current][node_no]

--fcc-editable-region--
```
