---
id: 687b45e16f38ad5f4e4bd799
title: Hatua 23
challengeType: 20
dashedName: step-23
---

# --description--

Sasa umeamua nodi inapaswa kuonyeshwa, unahitaji kuunda muundo wa njia yake ili iweze kuchapishwa kwa uwazi. Kwa hili utatumia usemi wa kizalishi.

Usemi wa kizalishi ni sawa na list comprehension, lakini badala ya kuunda orodha, huzalisha kila thamani moja kwa moja. Hutumia mabano ya kawaida `()` badala ya mabano ya mraba `[]`. Kwa mfano:

```py
numbers = [1, 2, 3]
squared = (x**2 for x in numbers)  # Generator expression
```

Ndani ya mzunguko baada ya tamko la `if`, tengeneza kigezo kinachoitwa `string_path`. Weka thamani ya usemi wa kizalishi unaobadilisha kila nambari ya nodi katika `paths[node_no]` kuwa mfuatano wa herufi kwa kutumia `str()`.

Usemi wa kizalishi unapaswa kurudia juu ya kila nambari ya nodi `n` katika `paths[node_no]`.

# --hints--

Unapaswa kutengeneza kigezo kinachoitwa `string_path`.

```js
({
    test: () => {
      assert(
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
for_loop.find_bodies()[0].has_variable("string_path")
      `)
      );
    }
});
```

Unapaswa kuweka thamani ya usemi wa kizalishi kwa `string_path`.

```js
({
  test: () => {
      runPython(`
      import ast
      assert isinstance(_Node(_code).find_function("shortest_path").find_for_loops()[1].find_bodies()[0].find_variable("string_path").tree.value, ast.GeneratorExp)
      `)
 }
})
```

Usemi wako wa kizalishi unapaswa kurudia juu ya `paths[node_no]`.

```js
({
    test: () => {
        runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert path_var.find_comp_iters()[0].is_equivalent("paths[node_no]")
      `);
    }
});
```

Usemi wako wa kizalishi unapaswa kutumia `n` kama kigezo cha kurudia juu ya `paths[node_no]`.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_targets()[0].is_equivalent("n")
      `);
    }
});
```

Usemi wako wa kizalishi unapaswa kutathmini `str(n)` kwa kila `n` katika `paths[node_no]`.

```js
({
  test: () => {
      runPython(`
func = _Node(_code).find_function("shortest_path")
for_loop = func.find_for_loops()[-1]
string_path_var = for_loop.find_bodies()[0].find_variable("string_path")
assert string_path_var.find_comp_expr().is_equivalent("str(n)")
      `);
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF] * n
    distances[start_node] = 0
    paths = [[node_no] for node_no in range(n)]
    visited = [False] * n
    
    for _ in range(n):
        min_distance = INF
        current = -1
        for node_no in range(n):
            if not visited[node_no] and distances[node_no] < min_distance:
                min_distance = distances[node_no]
                current = node_no
        
        if current == -1:
            break
        visited[current] = True
        
        for node_no in range(n):
            distance = matrix[current][node_no]
            if distance != INF and not visited[node_no]:
                new_distance = distances[current] + distance
                if new_distance < distances[node_no]:
                    distances[node_no] = new_distance
                    paths[node_no] = paths[current] + [node_no]
    
    targets = [target_node] if target_node is not None else range(n)
    for node_no in targets:
        if node_no == start_node or distances[node_no] == INF:
            continue
--fcc-editable-region--
        
--fcc-editable-region--
```
