---
id: 68e4c0ff9edcc92771d7549a
title: Hatua 6
challengeType: 20
dashedName: step-6
---

# --description--

Kwa sasa, orodha ya `distances` ina kipengele kimoja tu. Hata hivyo, unahitaji thamani moja ya umbali kwa kila nodi katika grafu.

Katika Python, unaweza kuzidisha orodha kwa nambari kamili ili kuirudia. Kwa mfano, `[0] * 3` huunda `[0, 0, 0]`.

Sahihisha orodha ya `distances` ili ianze na nakala `n` za `INF` kwa kuzidisha `[INF]` kwa `n`.

# --hints--

Unapaswa kusahihisha `distances` ili ianze na nakala `n` za `INF`.

```js
({
    test: () => {
      assert(
        runPython(
          `_Node(_code).find_function("shortest_path").find_variable("distances").is_equivalent("distances = [INF] * n")`
        )
      );
    }
});
```

# --seed--

## --seed-contents--

```py
INF = float('inf')
adj_matrix = [
    [0, 5, 3, INF, 11, INF],
    [5, 0, 1, INF, INF, 2],
    [3, 1, 0, 1, 5, INF],
    [INF, INF, 1, 0, 9, 3],
    [11, INF, 5, 9, 0, INF],
    [INF, 2, INF, 3, INF, 0],
]

--fcc-editable-region--
def shortest_path(matrix, start_node, target_node=None):
    n = len(matrix)
    distances = [INF]

--fcc-editable-region--
```
