---
id: 66b205e6eacba4c4e54ea434
title: 이메일 마스커 만들기
challengeType: 26
dashedName: build-an-email-masker
---

# --description--

이 실습에서는 이메일 주소의 사용자 이름 부분을 별표로 가릴 것입니다. 마스킹은 민감한 정보를 별표나 다른 문자로 숨기거나 대체하는 용어입니다.

예를 들어, 이메일 주소가 `myEmail@email.com`라면, 마스킹된 이메일 주소는 `m*****l@email.com`가 됩니다.

**목표:** 아래 사용자 스토리를 완성하고 모든 테스트를 통과하여 실습을 완료하세요.

**사용자 스토리:**

1. `maskEmail`를 인수로 받는 `email`라는 이름의 함수를 만드세요.
2. 함수 내부에서 `email`를 마스킹하고 도메인 이름을 덧붙이세요. `slice`, `repeat`, `indexOf` 또는 `replace` 같은 메서드를 사용해도 됩니다.
3. 함수 외부에 마스킹할 이메일 주소를 저장할 `email`라는 변수를 선언하세요.
4. `maskEmail` 변수를 인수로 하여 `email` 함수를 호출하고 결과를 콘솔에 출력하세요.
5. `maskEmail("apple.pie@example.com")`는 `"a*******e@example.com"`를 반환해야 합니다.
6. `maskEmail("freecodecamp@example.com")`는 `"f**********p@example.com"`를 반환해야 합니다.
7. `maskEmail("info@test.dev")`는 `"i**o@test.dev"`를 반환해야 합니다.
8. `maskEmail("user@domain.org")`는 `"u**r@domain.org"`를 반환해야 합니다.

# --hints--

`maskEmail`라는 이름의 함수를 정의하세요.

```js
assert.isFunction(maskEmail);
```

`maskEmail` 함수는 문자열 `email`를 인수로 받아야 합니다.

```js
assert.match(maskEmail.toString(), /\s*maskEmail\(\s*\w+\s*\)/);
```

함수 외부에 `email` 변수가 있어야 합니다.

```js
assert.isDefined(email);
```

`email` 변수에 유효한 이메일 주소를 할당하세요.

```js
assert.match(email, /^[a-zA-Z0-9._%-]+\+?[a-zA-Z0-9._%-]+@[a-zA-Z0-9.-]+\.[a-zA-Z]{2,}$/);
```

`maskEmail("apple.pie@example.com")`는 `"a*******e@example.com"`를 반환해야 합니다.

```js
assert.strictEqual(maskEmail("apple.pie@example.com"),"a*******e@example.com");
```

`maskEmail("freecodecamp@example.com")`는 `"f**********p@example.com"`를 반환해야 합니다.

```js
assert.strictEqual(maskEmail("freecodecamp@example.com"), "f**********p@example.com");
```

`maskEmail("info@test.dev")`는 `"i**o@test.dev"`를 반환해야 합니다.

```js
assert.strictEqual(maskEmail("info@test.dev"), "i**o@test.dev");
```

`maskEmail("user@domain.org")`는 `"u**r@domain.org"`를 반환해야 합니다.

```js
assert.strictEqual(maskEmail("user@domain.org"), "u**r@domain.org");
```

`maskEmail`는 올바른 결과를 출력해야 합니다.

```js
assert.strictEqual(maskEmail("johnappleseed@email.com"), "j***********d@email.com");
assert.strictEqual(maskEmail("john.doe@example.com"), "j******e@example.com");
assert.strictEqual(maskEmail("jane.smith@sample.org"), "j********h@sample.org");
assert.strictEqual(maskEmail("michael.jordan@testmail.net"), "m************n@testmail.net");
assert.strictEqual(maskEmail("mike+nospam@testmail.net"), "m*********m@testmail.net");
```

`maskEmail`를 `email` 인수로 호출한 결과를 로그에 출력하세요.

```js
assert.match(__helpers.removeJSComments(code), /console\.log\(maskEmail\(email\)\)/);
```

# --seed--

## --seed-contents--

```js

```

# --solutions--

```js
function maskEmail(email) {
  const atIndex = email.indexOf("@");
  const userName = email.slice(0, atIndex);
  const domain = email.slice(atIndex);

  const maskedName =
    userName[0] + "*".repeat(userName.length - 2) + userName[userName.length - 1] + domain;

  return maskedName;
}

const email = "apple.pie@example.com";
console.log(maskEmail(email));
```
