---
id: 68507cadbf36aa089a84ce2d
title: Implementar el algoritmo de Luhn
challengeType: 27
dashedName: implement-the-luhn-algorithm
---

# --description--

El algoritmo de Luhn, también conocido como el algoritmo "módulo 10" o "mod 10", es una fórmula de suma de verificación simple utilizada para validar una variedad de números de identificación, como los números de tarjetas de crédito. Estos son los pasos para validar un número usando el algoritmo de Luhn:

- Comenzando desde la derecha, y excluyendo el dígito más a la derecha (el dígito de verificación), duplica el valor de cada otro dígito.
- Si el resultado de duplicar un dígito es mayor que `9`, suma los dígitos para obtener un solo dígito. Alternativamente, puedes restar `9` del resultado.
- Toma la suma de todos los dígitos, incluyendo el dígito de verificación.
- Si la suma de todos los dígitos es un múltiplo de `10`, entonces el número es válido; de lo contrario, no es válido.

Digamos que tenemos el número `453914881`. Los pasos para validarlo usando el algoritmo de Luhn serían:

 ```md 
 Account number      4   5   3   9   1   4   8   8   1 
 Double every other  4  10   3  18   1   8   8  16   1 
 Sum 2-char digits   4   1   3   9   1   8   8   7   1 
 ```

- Luego suma todos los números, `4 + 1 + 3 + 9 + 1 + 8 + 8 + 7 + 1 = 42`.
- Dado que `42` no es un múltiplo de `10`, el número es inválido.

En este laboratorio, tú construirás un validador de tarjetas de crédito utilizando el algoritmo de Luhn.

Cumple con las historias de usuario a continuación y consigue que todas las pruebas pasen para completar el laboratorio.

**Historias de Usuario:**

1. Deberías definir una función llamada `verify_card_number` que tome una cadena de dígitos (que representa un número de tarjeta) y verifique si es válido según el algoritmo de Luhn.

2. Dentro de la función `verify_card_number`:

   - Deberías manejar cualquier guión o espacio que pueda estar presente en el número de tarjeta que se le pase.
   - Devuelve `¡VÁLIDO!` si el número de tarjeta es válido; de lo contrario, devuelve `¡INVÁLIDO!`.

Cuando completes el proyecto, deberías ver los siguientes mensajes dependiendo de la entrada:

| Número de Tarjeta   | Mensaje    |
| ------------------- | ---------- |
| 453914889           | ¡VÁLIDO!   |
| 4111-1111-1111-1111 | ¡VÁLIDO!   |
| 1234 5678 9012 3456 | ¡INVÁLIDO! |


# --hints--

Deberías tener una función llamada `verify_card_number`.

```js
({
    test: () => runPython(`
    assert _Node(_code).has_function('verify_card_number')
    `)
})
```

`verify_card_number('453914889')` debería devolver `¡VÁLIDO!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914889') == 'VALID!'
    `)
})
```

`verify_card_number('4111-1111-1111-1111')` debería devolver `¡VÁLIDO!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('4111-1111-1111-1111') == 'VALID!'
    `)
})
```

`verify_card_number('453914881')` debería devolver `¡INVÁLIDO!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('453914881') == 'INVALID!'
    `)
})
```

`verify_card_number('1234 5678 9012 3456')` debería devolver `¡INVÁLIDO!`.

```js
({
    test: () => runPython(`
    assert verify_card_number('1234 5678 9012 3456') == 'INVALID!'
    `)
})
```

`verify_card_number` debería devolver `¡VÁLIDO!` cuando se llame con un número de tarjeta de crédito válido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6467') == 'VALID!'
    assert verify_card_number('6011 1111 1111 1117') == 'VALID!'
    assert verify_card_number('3782-822463-10005') == 'VALID!'
    assert verify_card_number('5555 5555 5555 4444') == 'VALID!'
    assert verify_card_number('3056 9309 0259 04') == 'VALID!'
    assert verify_card_number('6011000990139424173') == 'VALID!'
    `)
})
```

`verify_card_number` debería devolver `¡INVÁLIDO!` cuando se llame con un número de tarjeta de crédito inválido.

```js
({
    test: () => runPython(`
    assert verify_card_number('4539 1488 0343 6466') == 'INVALID!'
    assert verify_card_number('6011 1111 1111 1116') == 'INVALID!'
    assert verify_card_number('3782 822463 10004') == 'INVALID!'
    assert verify_card_number('5555-5555-5555-4440') == 'INVALID!'
    assert verify_card_number('3056 9309 0259 05') == 'INVALID!'
    assert verify_card_number('6011000990139424175') == 'INVALID!'
    `)
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def verify_card_number(card_number):
    card_translation = str.maketrans({'-': '', ' ': ''})
    translated_card_number = card_number.translate(card_translation)

    sum_of_odd_digits = 0
    card_number_reversed = translated_card_number[::-1]
    odd_digits = card_number_reversed[::2]

    for digit in odd_digits:
        sum_of_odd_digits += int(digit)

    sum_of_even_digits = 0
    even_digits = card_number_reversed[1::2]
    for digit in even_digits:
        number = int(digit) * 2
        if number >= 10:
            number = (number // 10) + (number % 10)
        sum_of_even_digits += number

    total = sum_of_odd_digits + sum_of_even_digits
    if total % 10 == 0:
        return 'VALID!'
    else:
        return 'INVALID!'

# Example usage
verify_card_number('4111-1111-4555-1142')
```
