---
id: 69b2762b17d23fbee40a6c00
title: ステップ 2
challengeType: 1
dashedName: step-2
---

# --description--

マージソートアルゴリズムは次の3つの動作を行います。

- ソートされていないアイテムのシーケンスをサブパーツに分割する
- サブパーツ内のアイテムをソートする
- ソートされたサブパーツをマージする

これらは再帰的に行われ、サブパーツが完全なソート済みシーケンスにマージされるまで続きます。まずはシーケンスを分割することから始めましょう。

最初に、`mergeSort` 関数内で、`middlePoint` を使って `const` という変数を宣言し、シーケンスを再帰処理する前に半分に分割できるように `Math.floor(array.length / 2)` を代入してください。

# --hints--

`middlePoint` という変数を宣言し、`Math.floor(array.length / 2)` を代入する必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(const|let|var)\s+middlePoint\s*=\s*Math\.floor\(\s*array\.length\s*\/\s*2\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
}
```
