---
id: 69b279bae0f20eda63163cf4
title: ステップ 10
challengeType: 1
dashedName: step-10
---

# --description--

`if` の条件が `true` と評価される場合、それは `leftPart` 配列の要素が `rightPart` 配列の比較対象の要素よりも小さいことを意味します。

その場合、ソートされた配列の中で `leftArrayIndex` にある `leftPart` の要素を配置できます。

`if` ブロック内で、`leftPart[leftArrayIndex]` を `array[sortedIndex]` に代入してください。

# --hints--

`leftPart[leftArrayIndex]` の中で `array[sortedIndex]` を `if` に代入する必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /array\s*\[\s*sortedIndex\s*\]\s*=\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
