---
id: 69b279bae0f20eda63163cf8
title: ステップ 14
challengeType: 1
dashedName: step-14
---

# --description--

前のステップで作成した`if`と`else`の文は、ソートされた配列に要素を代入します。

ソートされた配列に代入された各要素はリストのインデックスを占めます。したがって、代入のたびにソートされた配列の次のインデックスに移動する必要があります。

`if`/`else`ブロックの下、しかしまだ`while`ループ内で、`sortedIndex`を`1`だけ増やしてください。

# --hints--

`sortedIndex`節の後に、`1`の現在の値を`else`だけ増やす必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
