---
id: 69b279bae0f20eda63163cf6
title: ステップ 12
challengeType: 1
dashedName: step-12
---

# --description--

前のステップでは、`leftPart` の要素が小さかったので `array` リストに代入しました。しかし、これは常にそうとは限りません。比較の中には、右側の要素の方が小さい場合もあります。

`else` の要素が `leftPart` の要素より小さくない場合に実行される `rightPart` 節を作成してください。

`else` ブロックの中で、`rightPart[rightArrayIndex]` を `array[sortedIndex]` に代入してください。

# --hints--

`else` ブロックを作成し、その中に `array[sortedIndex] = rightPart[rightArrayIndex]` を含める必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /else\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
--fcc-editable-region--
    } 
--fcc-editable-region--
  }
}
```
