---
id: 69b279bae0f20eda63163cf7
title: ステップ 13
challengeType: 1
dashedName: step-13
---

# --description--

`else` ブロック内で、`rightArrayIndex` を `1` だけ増やしてください。

# --hints--

ソート済み配列に要素を挿入した後で `rightArrayIndex` を増やす必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
--fcc-editable-region--
      
--fcc-editable-region--
    }
  }
}
```
