---
id: 69b2713ff6c21fed2d3ad88b
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

マージソートは、分割統治の原理を使ってデータの集合をソートするアルゴリズムです。つまり、集合をより小さな部分に「分割」し、それぞれの部分を独立してソートして「制覇」し、最後にソート済みの部分をマージします。

まず、`mergeSort` という名前の関数を作成し、パラメータ `array` を受け取るようにしてください。

# --hints--

`mergeSort` という名前の関数を、`array` という単一のパラメータで宣言してください。

```js
assert.isFunction(mergeSort);
const regex = __helpers.functionRegex('mergeSort', ['array']);
assert.match(__helpers.removeJSComments(code), regex);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
