---
id: 69b279bae0f20eda63163cfe
title: ステップ 20
challengeType: 1
dashedName: step-20
---

# --description--

`mergeSort()` 関数をテストする前に、`array` の長さが `1` 以下の場合に関数の実行を停止する基本ケースを作成する必要があります。

この基本ケースが再帰呼び出しをストップします。これがないと、リストがソート済みであっても、または要素がない場合でも、マージソートの処理が続行されてしまいます。

関数の宣言直後に、`if` が 1 以下かどうかをチェックする `array.length` 文を作成してください。`if` 文の本体内に、関数の実行を停止するための `return` 文を追加してください。

# --hints--

条件 `if` を持つ `array.length <= 1` 文を作成する必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{/
);
```

`return` 文の本体に `if` 文があるはずです。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*array\.length\s*<=\s*1\s*\)\s*\{[\s\S]*return\s*;?[\s\S]*\}/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
--fcc-editable-region--
  
--fcc-editable-region--
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }
  
  while (rightArrayIndex < rightPart.length) {
    array[sortedIndex] = rightPart[rightArrayIndex];
    rightArrayIndex += 1;
    sortedIndex += 1;
  }
}
```
