---
id: 69b279bae0f20eda63163cef
title: ステップ 5
challengeType: 1
dashedName: step-5
---

# --description--

`array` リストを2つの別々のリストに分割したので、今度はそれぞれのリストをさらに分割して、すべての要素が単独のリストになるまで続けます。1つの数字だけのリストは常にソート済みです。

そのために、関数の中で再帰的に `mergeSort` を呼び出し、呼び出しの引数として `leftPart` を渡してください。

# --hints--

同じ `mergeSort(leftPart)` 関数の中で `mergeSort` を呼び出す必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*leftPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
