---
id: 69b278519ef46b2fd8bf1546
title: ステップ 3
challengeType: 1
dashedName: step-3
---

# --description--

すでに `middlePoint` の値があります。これを使って配列を2つの半分に分割し、左側の部分をソートの準備として保持してください。

スライス構文を使って `array` の左半分を抽出し、`leftPart` という名前の変数に代入してください。

# --hints--

`leftPart` 関数の中で `mergeSort` 変数を宣言してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftPart\s*=/
);
```

`array.slice(0, middlePoint)` に `leftPart` を代入してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /leftPart\s*=\s*array\.slice\(\s*0\s*,\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
