---
id: 69b279bae0f20eda63163cf2
title: ステップ 8
challengeType: 1
dashedName: step-8
---

# --description--

次に、`leftPart` と `rightPart` の両方に要素が残っている限り続くループを作成する必要があります。

そのために、2つの条件を持つ `while` ループを作成してください。1つは `leftArrayIndex` が `leftPart` の長さより小さいかどうかをチェックし、もう1つは `rightArrayIndex` が `rightPart` の長さより小さいかどうかをチェックします。

# --hints--

`while` ループを作成し、`leftArrayIndex` が `leftPart` の長さより小さいかどうか、そして `rightArrayIndex` が `rightPart` の長さより小さいかどうかをチェックしてください。

```js
assert.match(
  __helpers.removeJSComments(code),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*&&\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
