---
id: 69b279bae0f20eda63163cf1
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

いよいよリスト（`leftPart` と `rightPart`）をソートして元の `array` にマージする時です。

両方のリストの要素を比較し、小さい方の要素をメインのリストにマージすることでこれを行います。`leftPart` と `rightPart` のすべてのインデックスでこの比較を行います。

3つの変数 `leftArrayIndex`、`rightArrayIndex`、`sortedIndex` を作成し、それぞれに `0` を設定してください。これらの変数はソートの過程で各インデックスを追跡するのに役立ちます。

# --hints--

`leftArrayIndex` という変数を宣言し、`0` で初期化してください。

```js
const explorer = await __helpers.Explorer(code);
const { leftArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  leftArrayIndex,
  'You should declare a leftArrayIndex variable inside mergeSort'
);
assert.isTrue(
  leftArrayIndex.value.matches('0'),
  'leftArrayIndex should be initialized to 0'
);
```

`rightArrayIndex` という変数を宣言し、`0` で初期化してください。

```js
const explorer = await __helpers.Explorer(code);
const { rightArrayIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  rightArrayIndex,
  'You should declare a rightArrayIndex variable inside mergeSort'
);
assert.isTrue(
  rightArrayIndex.value.matches('0'),
  'rightArrayIndex should be initialized to 0'
);
```

`sortedIndex` という変数を宣言し、`0` で初期化してください。

```js
const explorer = await __helpers.Explorer(code);
const { sortedIndex } = explorer.allFunctions.mergeSort.variables;
assert.exists(
  sortedIndex,
  'You should declare a sortedIndex variable inside mergeSort'
);
assert.isTrue(
  sortedIndex.value.matches('0'),
  'sortedIndex should be initialized to 0'
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
--fcc-editable-region--
  
--fcc-editable-region--
}
```
