---
id: 69b278d82f0e514f32acb724
title: ステップ 4
challengeType: 1
dashedName: step-4
---

# --description--

左半分をスライスした後、マージのステップを完了するには右側も必要です。

スライス構文を使って `array` の右半分を抽出し、`rightPart` という名前の変数に代入してください。

# --hints--

`rightPart` 関数の中に `mergeSort` という変数があるはずです。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightPart\s*=/
);
```

配列の右半分をスライス構文で抽出し、`rightPart` 変数に代入してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /rightPart\s*=\s*array\.slice\(\s*middlePoint\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
