---
id: 69b279bae0f20eda63163cfd
title: ステップ 19
challengeType: 1
dashedName: step-19
---

# --description--

最後に、右側の残りの値をすべて `array` にコピーしてください。

そのために、`while` が `rightArrayIndex` より小さい間ループする `rightPart.length` を追加し、`rightPart[rightArrayIndex]` を `array[sortedIndex]` に代入し、`rightArrayIndex` と `sortedIndex` の両方をインクリメントしてください。

# --hints--

`while` の条件で `rightArrayIndex < rightPart.length` ループを作成してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{/
);
```

ループの本体では `rightPart[rightArrayIndex]` を `array[sortedIndex]` に移動してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*array\s*\[\s*sortedIndex\s*\]\s*=\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*;?/
);
```

代入の後、`rightArrayIndex` と `sortedIndex` の両方をインクリメントしてください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*rightArrayIndex\s*<\s*rightPart\.length\s*\)\s*\{[\s\S]*rightArrayIndex\s*(\+\+|\+=\s*1)\s*;?[\s\S]*sortedIndex\s*(\+\+|\+=\s*1)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

  while (leftArrayIndex < leftPart.length) {
    array[sortedIndex] = leftPart[leftArrayIndex];
    leftArrayIndex += 1;
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
