---
id: 69b279bae0f20eda63163cf1
title: ステップ 7
challengeType: 1
dashedName: step-7
---

# --description--

いよいよリスト（`leftPart` と `rightPart`）をソートして元の `array` にマージする時です。

両方のリストの要素を比較し、小さい方の要素をメインのリストにマージすることでこれを行います。`leftPart` と `rightPart` のすべてのインデックスでこの比較を行います。

3つの変数 `leftArrayIndex`、`rightArrayIndex`、`sortedIndex` を作成し、それぞれに `0` を設定してください。これらの変数はソートの過程で各インデックスを追跡するのに役立ちます。

# --hints--

`leftArrayIndex` という変数を宣言し、`0` で初期化してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+leftArrayIndex\s*=\s*0\s*;?/
);
```

`rightArrayIndex` という変数を宣言し、`0` で初期化してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+rightArrayIndex\s*=\s*0\s*;?/
);
```

`sortedIndex` という変数を宣言し、`0` で初期化してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /(let|const|var)\s+sortedIndex\s*=\s*0\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);

--fcc-editable-region--
  
--fcc-editable-region--
}
```
