---
id: 69b279bae0f20eda63163cf0
title: ステップ 6
challengeType: 1
dashedName: step-6
---

# --description--

関数本体の一番下で、`mergeSort()` 関数をもう一度呼び出してください。今回は、`rightPart` を引数として関数呼び出しに渡してください。

# --hints--

関数本体の一番下で、`mergeSort()` を引数 `rightPart` で呼び出す必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /mergeSort\s*\(\s*rightPart\s*\)\s*;?/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
--fcc-editable-region--
  
--fcc-editable-region--
}
```
