---
id: 69b279bae0f20eda63163cf3
title: ステップ 9
challengeType: 1
dashedName: step-9
---

# --description--

ループの中で、それぞれの半分から現在の要素を比較し始めて、小さい方の値を先にマージできるようにしてください。

`if` ブロックを追加して、`leftPart` の現在の要素が `rightPart` の現在の要素より小さいかどうかをチェックしてください。 
# --hints--

`if` の条件は `leftPart[leftArrayIndex]` が `rightPart[rightArrayIndex]` より小さいかどうかをチェックする必要があります。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /if\s*\(\s*leftPart\s*\[\s*leftArrayIndex\s*\]\s*<\s*rightPart\s*\[\s*rightArrayIndex\s*\]\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;
  
  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
--fcc-editable-region--
    
--fcc-editable-region--
  }
}
```
