---
id: 69b2713ff6c21fed2d3ad88b
title: ステップ 1
challengeType: 1
dashedName: step-1
---

# --description--

マージソートは、分割統治の原理を使ってデータの集合をソートするアルゴリズムです。つまり、集合をより小さな部分に「分割」し、それぞれの部分を独立してソートして「制覇」し、最後にソート済みの部分をマージします。

まず、`mergeSort` という名前の関数を作成し、パラメータ `array` を受け取るようにしてください。

# --hints--

`mergeSort` という名前の関数を、`array` という単一のパラメータで宣言してください。

```js
assert.isFunction(mergeSort);

const explorer = await __helpers.Explorer(code);
const mergeSortFn = explorer.allFunctions.mergeSort;
assert.exists(mergeSortFn, 'You should declare a mergeSort function');

const params = mergeSortFn.parameters;
assert.lengthOf(
  params,
  1,
  'mergeSort should take a single parameter'
);
assert.isTrue(
  params[0].matches('array'),
  'The parameter of mergeSort should be named array'
);
```

# --seed--

## --seed-contents--

```js
--fcc-editable-region--

--fcc-editable-region--
```
