---
id: 69b279bae0f20eda63163cf9
title: ステップ 15
challengeType: 1
dashedName: step-15
---

# --description--

作成した`while`ループは、`leftPart`の要素と`rightPart`の要素を比較し、小さい方の要素をメインの`array`リストに追加します。

この操作は比較する要素がなくなるまで続きます。しかし、`leftPart`にはまだ要素が残っている場合があり、`rightPart`には要素がない場合もありますし、その逆もあります。

`while`が`leftArrayIndex`の長さより小さいときに実行される別の`leftPart`ループを作成してください。次のステップで、`leftPart`に残っている要素を`array`リストにコピーするために使います。

# --hints--

`while`という条件で`leftArrayIndex < leftPart.length`ループを作成してください。

```js
assert.match(
  __helpers.removeJSComments(mergeSort.toString()),
  /while\s*\(\s*leftArrayIndex\s*<\s*leftPart\.length\s*\)\s*\{/
);
```

# --seed--

## --seed-contents--

```js
function mergeSort(array) {
  const middlePoint = Math.floor(array.length / 2);
  const leftPart = array.slice(0, middlePoint);
  const rightPart = array.slice(middlePoint);
  
  mergeSort(leftPart);
  mergeSort(rightPart);
  
  let leftArrayIndex = 0;
  let rightArrayIndex = 0;
  let sortedIndex = 0;

  while (leftArrayIndex < leftPart.length && rightArrayIndex < rightPart.length) {
    if (leftPart[leftArrayIndex] < rightPart[rightArrayIndex]) {
      array[sortedIndex] = leftPart[leftArrayIndex];
      leftArrayIndex += 1;
    } else {
      array[sortedIndex] = rightPart[rightArrayIndex];
      rightArrayIndex += 1;
    }
    sortedIndex += 1;
  }

--fcc-editable-region--
  
--fcc-editable-region--
}
```
