---
id: 5dc24073f86c76b9248c6ebb
title: Step 8
challengeType: 0
dashedName: step-8
---

# --description--

HTML <dfn>attributes</dfn> are special words used inside the opening tag of an element to control the element's behavior. The `src` attribute in an `img` element specifies the image's URL (where the image is located).

Ecco un esempio di un elemento `img` con un attributo `src` che punta al logo di freeCodeCamp:

```html
<img src="https://cdn.freecodecamp.org/platform/universal/fcc_secondary.svg">
```

Inside the existing `img` element, add a `src` attribute with this URL:

`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`

# --hints--

Il codice dovrebbe avere un elemento `img`. Potresti aver rimosso l'elemento `img` o non hai circondato il valore dell'attributo `src` con le virgolette.

```js
assert(document.querySelector('img'));
```

Your `img` element should have a `src` attribute. Hai omesso l'attributo o hai un refuso. Assicurati che ci sia uno spazio tra il nome dell'elemento e il nome dell'attributo.

```js
assert(document.querySelector('img').src);
```

L'attributo `src` dell'elemento `img` dovrebbe essere '`https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg`'. Hai omesso l'URL o hai un refuso. Le maiuscole/minuscole nell'URL sono importanti.

```js
assert.strictEqual(document.querySelector('img').src, 'https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg');
```

Anche se hai impostato l'URL corretto per l'`src` dell'elemento `img`, è raccomandato circondare sempre il valore di un attributo con le virgolette.

```js
assert.notMatch(code, /\<img\s+src\s*=\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg/);
```

Your `img` element should be closed either with `>` or `/>`.

```js
assert.match(code, /<img\s+src\s*=\s*("|')\s*https:\/\/cdn\.freecodecamp\.org\/curriculum\/cat-photo-app\/relaxing-cat\.jpg\s*\1\s*\/?>/);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <h2>Cat Photos</h2>
      <!-- TODO: Add link to cat photos -->
      <p>Everyone loves cute cats online!</p>
--fcc-editable-region--
      <img>
--fcc-editable-region--
    </main>
  </body>
</html>
```

