---
id: 5ef9b03c81a63668521804d2
title: Step 29
challengeType: 0
dashedName: step-29
---

# --description--

The code for an ordered list (`ol`) is similar to an unordered list, but list items in an ordered list are numbered when displayed.

Below the `h3` element, add an ordered list with these three list items:

`flea treatment` `thunder` `other cats`

# --hints--

L'elemento `ol` dovrebbe avere un tag di apertura. I tag di apertura hanno questa sintassi: `<elementName>`.

```js
assert(document.querySelector('ol'));
```

L'elemento `ol` dovrebbe avere un tag di chiusura. I tag di chiusura hanno un carattere `/` subito dopo il carattere `<`.

```js
assert(code.match(/<\/ol>/));
```

Ci dovrebbe essere un elemento `ol` proprio sopra il tag di chiusura del secondo elemento `section`. Sono nell'ordine sbagliato.

```js
assert.equal(document.querySelectorAll('main > section')?.[1]?.lastElementChild?.nodeName, 'OL');
```

I tre elementi `li` dovrebbero essere annidati all’interno dell’elemento `ol`.

```js
assert(
  [...document.querySelectorAll('li')].filter(
    (item) => item.parentNode.nodeName === 'OL'
  ).length === 3
);
```

Dovresti avere tre elementi `li` con il testo `flea treatment`, `thunder` e `other cats` in qualsiasi ordine.

```js
assert.deepStrictEqual(
  [...document.querySelectorAll('li')]
    .filter((item) => item.parentNode.nodeName === 'OL')
    .map((item) => item.innerText.toLowerCase())
    .sort((a, b) => a.localeCompare(b)),
  ['flea treatment', 'other cats', 'thunder']
);
```

Dovresti avere un solo elemento `ol`.

```js 
assert([...document.querySelectorAll('ol')].length == 1);
```

# --seed--

## --seed-contents--

```html
<html>
  <body>
    <main>
      <h1>CatPhotoApp</h1>
      <section>
        <h2>Cat Photos</h2>
        <p>Everyone loves <a href="https://cdn.freecodecamp.org/curriculum/cat-photo-app/running-cats.jpg">cute cats</a> online!</p>
        <p>See more <a target="_blank" href="https://freecatphotoapp.com">cat photos</a> in our gallery.</p>
        <a href="https://freecatphotoapp.com"><img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/relaxing-cat.jpg" alt="A cute orange cat lying on its back."></a>
      </section>
      <section>
        <h2>Cat Lists</h2>
        <h3>Things cats love:</h3>
        <ul>
          <li>catnip</li>
          <li>laser pointers</li>
          <li>lasagna</li>
        </ul>
        <figure>
          <img src="https://cdn.freecodecamp.org/curriculum/cat-photo-app/lasagna.jpg" alt="A slice of lasagna on a plate.">
          <figcaption>Cats <em>love</em> lasagna.</figcaption>  
        </figure>
--fcc-editable-region--
        <h3>Top 3 things cats hate:</h3>

--fcc-editable-region--
      </section>
    </main>

  </body>
</html>
```

