---
id: 695774002591bbc5f8cf3e53
title: Eine Funktion zum Anwenden von Rabatten erstellen
challengeType: 27
dashedName: build-a-discount-calculator
---

# --description--

In diesem Labor schreiben Sie eine Funktion, die den Endpreis eines Artikels nach Anwendung eines prozentualen Rabatts berechnet.

Zum Beispiel, wenn der Preis eines Artikels `50` beträgt und ein Rabatt von `20` angewendet wird, beträgt der Rabattbetrag `10` und der Endpreis `40`.

**Ziel:** Erfüllen Sie die untenstehenden User Stories, um das Labor abzuschließen.

**User Stories:**

1. Sie sollten eine Funktion mit dem Namen `apply_discount` definieren.

2. Die `apply_discount`-Funktion sollte genau zwei Parameter annehmen: `price` und `discount`.

3. Wenn `price` keine Zahl ist (`int` oder `float`), sollte die Funktion den String `The price should be a number` zurückgeben.

4. Wenn `discount` keine Zahl ist (`int` oder `float`), sollte die Funktion den String `The discount should be a number` zurückgeben.

5. Wenn `price` kleiner oder gleich `0` ist, sollte die Funktion den String `The price should be greater than 0` zurückgeben.

6. Wenn `discount` kleiner als `0` oder größer als `100` ist, sollte die Funktion den String `The discount should be between 0 and 100` zurückgeben.

7. Wenn beide Eingaben gültig sind, sollte die Funktion den Rabatt als Prozentsatz des Preises berechnen.

8. Die Funktion sollte den Endpreis nach Anwendung des Rabatts zurückgeben.

# --hints--

Sie sollten eine Funktion mit dem Namen `apply_discount` haben.

```js
({
  test: () => runPython(
    `assert _Node(_code).has_function("apply_discount")`
  )
})
```

Ihre `apply_discount`-Funktion sollte zwei Parameter annehmen: `price` und `discount`.

```js
({
  test: () => runPython(
`
func = _Node(_code).find_function("apply_discount")

assert func.has_args("price, discount")
`
  )
})
```

Wenn `apply_discount` mit einem `price` (erstes Argument) aufgerufen wird, das keine Zahl ist (`int` oder `float`), sollte sie `The price should be a number` zurückgeben.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "100", [1,2], {"a":1}, (1,2), None]
for v in wrong_values:
    assert apply_discount(v, 10) == "The price should be a number"
`
  )
})
```

Wenn `apply_discount` mit einem `discount` (zweites Argument) aufgerufen wird, das keine Zahl ist (`int` oder `float`), sollte sie `The discount should be a number` zurückgeben.

```js
({
  test: () => runPython(
`
wrong_values = ["HelloWorld", "10", [1], {"a":1}, (1,), None]
for v in wrong_values:
    assert apply_discount(100, v) == "The discount should be a number"
`
  )
})
```

Wenn `apply_discount` mit einem `price` aufgerufen wird, der kleiner oder gleich `0` ist, sollte sie `The price should be greater than 0` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(0, 10) == "The price should be greater than 0"
assert apply_discount(-5, 10) == "The price should be greater than 0"
`
  )
})
```

Wenn `apply_discount` mit einem `discount` aufgerufen wird, der kleiner als `0` oder größer als `100` ist, sollte sie `The discount should be between 0 and 100` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(100, -1) == "The discount should be between 0 and 100"
assert apply_discount(100, 101) == "The discount should be between 0 and 100"
`
  )
})
```

`apply_discount(100, 20)` sollte `80` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(100, 20) == 80
`
  )
})
```

`apply_discount(200, 50)` sollte `100` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(200, 50) == 100
`
  )
})
```

`apply_discount(50, 0)` sollte `50` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(50, 0) == 50
`
  )
})
```

Wenn `apply_discount` mit einem Rabatt von `100` aufgerufen wird, sollte sie `0` zurückgeben.

```js
({
  test: () => runPython(
`
assert apply_discount(120, 100) == 0
`
  )
})
```

`apply_discount(74.5, 20.0)` sollte `59.6` zurückgeben.

```js
({
  test: () => runPython(
`
result = apply_discount(74.5, 20.0)
assert result == 59.6
`
  )
})
```

# --seed--

## --seed-contents--

```py

```

# --solutions--

```py
def apply_discount(price, discount):
    # Validation steps
    if not isinstance(price, (int, float)):
        return "The price should be a number"
    if not isinstance(discount, (int, float)):
        return "The discount should be a number"
    if price <= 0:
        return "The price should be greater than 0"
    if discount < 0 or discount > 100:
        return "The discount should be between 0 and 100"

    discount_amount = price * discount / 100
    final_price = price - discount_amount
    return final_price
```
