---
id: 68e524c94e2e4ab7b4bfc7e2
title: ステップ 1
challengeType: 20
dashedName: step-1
---

# --description--

このワークショップでは、さまざまな割引戦略を商品に適用できる割引計算機を作成します。システムは複数の割引ルールに基づいて顧客にとって最適な価格を決定します。

まず、`Product`という名前のクラスを定義してください。クラスに`__init__`メソッドを作成し、パラメータとして`self`、`name`、`price`を指定してください。`__init__`メソッドの中で、`name`と`price`のパラメータを同じ名前のインスタンス属性に代入してください。

# --hints--

`Product`という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("Product")`)
})
```

`Product`クラスには`__init__`メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__init__")`)
})
```

`__init__`メソッドには、`self`、`name`、`price`というパラメータが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name, price")`)
})
```

`name`を`self.name`に`__init__`メソッド内で代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.name = name")`)
})
```

`price`を`self.price`に`__init__`メソッド内で代入してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_stmt("self.price = price")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
```
