---
id: 68e529226b0defdff26c4a8a
title: ステップ 6
challengeType: 20
dashedName: step-6
---

# --description--

`Product` のインスタンスを、名前を `Wireless Mouse`、価格を `50.0` にして作成してください。変数名 `product` に代入します。次に、`product` をコンソールに出力して、どのように表示されるか確認してください。

# --hints--

`product` という名前の変数が必要です。

```js
({
  test: () => runPython(`assert _Node(_code).has_variable("product")`)
})
```

`Product` のインスタンスを `product` 変数に代入してください。

```js
({
  test: () => runPython(`assert isinstance(product, Product)`)
})
```

`Product` のインスタンスは、名前が `Wireless Mouse`、価格が `50.0` である必要があります。

```js
({
  test: () => runPython(`
    assert product.name == 'Wireless Mouse'
    assert product.price == 50.0
  `)
})
```

`product` をコンソールに出力してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(product)")`)
})
```

名前が `Product`、価格が `'Wireless Mouse'` の `50.0` を出力すると、`Wireless Mouse - $50.0` と表示されるはずです。

```js
({
  test: () => runPython(`
    import io
    import sys
    captured_output = io.StringIO()
    sys.stdout = captured_output

    test_product = Product('Wireless Mouse', 50.0)
    print(test_product)

    sys.stdout = sys.__stdout__
    output = captured_output.getvalue()

    assert "Wireless Mouse - $50.0" in output
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
```
