---
id: 68e5293bd00d2fe134f58986
title: ステップ 14
challengeType: 20
dashedName: step-14
---

# --description--

今度は、`apply_discount` メソッドを `PercentageDiscount` に実装してください。

このメソッドは割引価格を計算して返す必要があります。

計算式は `product.price * (1 - self.percent / 100)` です。

# --hints--

あなたの `PercentageDiscount` クラスには `apply_discount` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("apply_discount")`)
})
```

あなたの `apply_discount` メソッドにはパラメータ `self` と `product: Product` が必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_args("self, product: Product")`)
})
```

あなたの `apply_discount` メソッドには戻り値の型ヒントとして `float` を付けてください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("apply_discount").has_returns("float")`)
})
```

あなたの `apply_discount` メソッドは計算式 `product.price * (1 - self.percent / 100)` を使って割引価格を計算する必要があります。

```js
({
  test: () => runPython(`
    test_product = Product("Test", 100.0)
    discount = PercentageDiscount(20)
    result = discount.apply_discount(test_product)
    assert result == 80.0
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
