---
id: 68e52994bc7ea2e3dec530a1
title: ステップ 23
challengeType: 20
dashedName: step-23
---

# --description--

複数の割引戦略を管理し、最適な価格を見つける方法が必要です。

`DiscountEngine` という名前のクラスを作成してください。このクラスはすべての割引戦略を管理し、商品の最適な価格を計算します。

`__init__` メソッドを用意し、`self` と `strategies` をパラメータとして受け取ります。ここで `strategies` は `list[DiscountStrategy]` の型ヒントを持ちます。`list[DiscountStrategy]` を使うことで、パラメータが `DiscountStrategy` オブジェクトを含むリストであることを指定します。メソッドの戻り値の型は `None` にしてください。`__init__` メソッドの中で、`strategies` をインスタンス属性として代入してください。

# --hints--

`DiscountEngine` という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountEngine")`)
})
```

`DiscountEngine` クラスには `__init__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("__init__")`)
})
```

`__init__` メソッドには `self` と `strategies: list[DiscountStrategy]` のパラメータが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_args("self, strategies: list[DiscountStrategy]")`)
})
```

`__init__` メソッドの戻り値の型は `None` にしてください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_returns("None")`)
})
```

`strategies` メソッドの中で `self.strategies` を `__init__` に代入してください。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("DiscountEngine").find_function("__init__").has_stmt("self.strategies = strategies")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
