---
id: 68e529226b0defdff26c4a8b
title: ステップ 7
challengeType: 20
dashedName: step-7
---

# --description--

これから割引システムのインターフェースを作り始めます。Pythonでは、他のクラスの設計図となる抽象クラスを定義する方法が用意されています。これは`ABC`モジュールの`abc`クラス（抽象基底クラス）を使って行います。

コードの先頭に、次のインポート文を追加してください。

```py
from abc import ABC, abstractmethod
```

これにより、`ABC`クラスと`abstractmethod`デコレーターがインポートされます。これらは抽象クラスを継承するすべてのクラスで実装が必須となるメソッドを定義するために使います。

# --hints--

コードの先頭に、上記のインポート文があることを確認してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_imports()[0].is_equivalent("from abc import ABC, abstractmethod")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--

--fcc-editable-region--
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

product = Product('Wireless Mouse', 50.0)
print(product)
```
