---
id: 68e5293bd00d2fe134f58982
title: ステップ 9
challengeType: 20
dashedName: step-9
---

# --description--

前のレッスンで説明したように、抽象メソッドは抽象基底クラスを継承するすべてのクラスが実装しなければならないメソッドです。これらは契約として機能し、すべての割引戦略が特定の必須メソッドを持つことを保証します。

`DiscountStrategy` クラスの中に、`is_applicable` という名前のメソッドを定義してください。メソッド定義の上に `@abstractmethod` デコレーターを使います。メソッドは次のパラメータを持つべきです：`self`、`product`、および `user_tier`。メソッド本体には `pass` 文を使ってください。

メソッドのシグネチャを明確にするために型ヒントを追加してください：`product` は `Product` 型、`user_tier` は `str` 型であるべきです。


# --hints--

あなたの `DiscountStrategy` クラスには `is_applicable` という名前のメソッドがあるべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("is_applicable")`)
})
```

あなたの `is_applicable` メソッドには `@abstractmethod` デコレーターがあるべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_decorators("abstractmethod")`)
})
```

あなたの `is_applicable` メソッドには `self`、`product: Product`、および `user_tier: str` のパラメータがあるべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

--fcc-editable-region--
class DiscountStrategy(ABC):
    pass
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
