---
id: 68e529226b0defdff26c4a89
title: ステップ 3
challengeType: 20
dashedName: step-3
---

# --description--

今、`Product` のインスタンスを作成して出力すると、あまり情報が得られないデフォルトのオブジェクト表現が表示されます。例えば、`<__main__.Product object at 0x7f9b8c2d1d30>` のように見えるかもしれません。

デフォルトのオブジェクト表現を読みやすい文字列に変換するには、`__str__` クラスに `Product` メソッドを作成し、`name - $price` と `name` を対応する属性の値に置き換えた形式 `price` の文字列を返してください。

# --hints--

`Product` クラスには `__str__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").has_function("__str__")`)
})
```

`__str__` メソッドは `name - $price` の形式の文字列を返す必要があります。

```js
({
  test: () => runPython(`
    test_product = Product("Test Item", 25.5)
    assert str(test_product) == "Test Item - $25.5"
  `)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float):
        self.name = name
        self.price = price
--fcc-editable-region--
    
--fcc-editable-region--
```
