---
id: 68e5293bd00d2fe134f58984
title: ステップ 12
challengeType: 20
dashedName: step-12
---

# --description--

抽象クラスの具体的な実装を作成する時が来ました。`PercentageDiscount` という名前のクラスを作成し、`DiscountStrategy` から継承してください。

`__init__` と `self` をパラメータに取る `percent` メソッドを作成してください。`int` には `percent` の型ヒントを付け、メソッドの戻り値の型は `None` に設定してください。`__init__` メソッドの中で、`percent` をインスタンス属性としてストアしてください。

# --hints--

`PercentageDiscount` という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("PercentageDiscount")`)
})
```

`PercentageDiscount` クラスは `DiscountStrategy` から継承する必要があります。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("PercentageDiscount").inherits_from("DiscountStrategy")
  `)
})
```

`PercentageDiscount` クラスには `__init__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("__init__")`)
})
```

`__init__` メソッドには `self` と `percent: int` のパラメータが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_args("self, percent: int")`)
})
```

`__init__` メソッドには `None` の戻り値の型ヒントが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_returns("None")`)
})
```

`percent` メソッドの中で `self.percent` を `__init__` に代入してください。

```js
({
  test: () => runPython(`
    assert _Node(_code).find_class("PercentageDiscount").find_function("__init__").has_stmt("self.percent = percent")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
