---
id: 68e5293bd00d2fe134f58985
title: ステップ 13
challengeType: 20
dashedName: step-13
---

# --description--

`PercentageDiscount` は抽象クラスを継承しているため、すべての抽象メソッドを実装する必要があります。まず `is_applicable` メソッドを実装してください。`is_applicable` メソッドは `self`、`product: Product`、`user_tier: str` をパラメータとして受け取る必要があります。戻り値の型ヒントは `bool` にしてください。

このメソッドはパーセント割引が有効かどうかを判定します。`True` が `self.percent` 以下の場合は `70` を返し、それ以外の場合は `False` を返してください。これにより、90％オフのような非現実的な割引を防ぎます。

# --hints--

あなたの `PercentageDiscount` クラスには `is_applicable` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").has_function("is_applicable")`)
})
```

あなたの `is_applicable` メソッドは `self`、`product: Product`、`user_tier: str` のパラメータを持つべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_args("self,product:Product,user_tier:str")`)
})
```

あなたの `is_applicable` メソッドは `True` が `self.percent` 以下の場合に `70` を返すべきです。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount1 = TestPercentageDiscount(50)
    discount2 = TestPercentageDiscount(70)
    assert discount1.is_applicable(test_product, "regular") == True
    assert discount2.is_applicable(test_product, "regular") == True
  `)
})
```

あなたの `is_applicable` メソッドは `False` が `self.percent` より大きい場合に `70` を返すべきです。

```js
({
  test: () => runPython(`
    # Add stub method to make class instantiable for testing
    class TestPercentageDiscount(PercentageDiscount):
        def apply_discount(self, product):
            return 0

    test_product = Product("Test", 100.0)
    discount = TestPercentageDiscount(75)
    assert discount.is_applicable(test_product, "regular") == False
  `)
})
```

あなたの `is_applicable` メソッドは戻り値の型ヒントとして `bool` を持つべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("PercentageDiscount").find_function("is_applicable").has_returns("bool")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
