---
id: 68e52994bc7ea2e3dec530a2
title: ステップ 24
challengeType: 20
dashedName: step-24
---

# --description--

`calculate_best_price` クラスに `DiscountEngine` という名前のメソッドを作成してください。このメソッドは `self`、`product`、`user_tier` をパラメータとして受け取り、`float` を返す必要があります。

今はメソッドの本体で `pass` を使うだけで構いません。ロジックは次のステップで実装します。

# --hints--

`DiscountEngine` クラスには `calculate_best_price` という名前のメソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").has_function("calculate_best_price")`)
})
```

`calculate_best_price` メソッドには `self`、`product: Product`、`user_tier: str` のパラメータが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").has_args("self, product:Product, user_tier:str")`)
})
```

`calculate_best_price` メソッドには `-> float` の戻り値の型ヒントが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountEngine").find_function("calculate_best_price").has_returns("float")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

class FixedAmountDiscount(DiscountStrategy):
    def __init__(self, amount: int) -> None:
        self.amount = amount

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return product.price * 0.9 > self.amount

    def apply_discount(self, product: Product) -> float:
        return product.price - self.amount

class PremiumUserDiscount(DiscountStrategy):
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return user_tier.lower() == 'premium'

    def apply_discount(self, product: Product) -> float:
        return product.price * 0.8

class DiscountEngine:
    def __init__(self, strategies: list[DiscountStrategy]) -> None:
        self.strategies = strategies
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))

fixed_discount = FixedAmountDiscount(5)
print(fixed_discount.apply_discount(product))
```
