---
id: 68e5293bd00d2fe134f58987
title: ステップ 16
challengeType: 20
dashedName: step-16
---

# --description--

別の割引戦略を作成します。`FixedAmountDiscount` という名前のクラスを作成し、`DiscountStrategy` から継承してください。

`__init__` と `self` をパラメータに取る `amount` メソッドを作成してください。`amount` パラメータには `int` の型ヒントを付け、メソッドの戻り値には `None` の型ヒントを付けてください。`__init__` メソッドの中で、`amount` をインスタンス属性としてストアしてください。

# --hints--

`FixedAmountDiscount` という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("FixedAmountDiscount")`)
})
```

`FixedAmountDiscount` クラスは `DiscountStrategy` から継承する必要があります。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").inherits_from("DiscountStrategy")`)
})
```

`FixedAmountDiscount` クラスには `__init__` メソッドが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").has_function("__init__")`)
})
```

`__init__` メソッドには `self` と `amount: int` のパラメータが必要です。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_args("self, amount: int")`)
})
```

`__init__` メソッドの戻り値の型は `None` である必要があります。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_returns("None")`)
})
```

`amount` メソッドの中で `self.amount` を `__init__` に代入してください。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("FixedAmountDiscount").find_function("__init__").has_stmt("self.amount = amount")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70

    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

discount = PercentageDiscount(10)
print(discount.apply_discount(product))
```
