---
id: 68e5293bd00d2fe134f58981
title: ステップ 8
challengeType: 20
dashedName: step-8
---

# --description--

`DiscountStrategy` から継承する `ABC` という名前のクラスを作成してください。これはすべての割引戦略のインターフェースを定義する抽象基底クラスになります。

今のところ、クラス本体には `pass` を使って空のクラスを作成してください。

# --hints--

`DiscountStrategy` という名前のクラスを作成してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_class("DiscountStrategy")`)
})
```

`DiscountStrategy` クラスは `ABC` から継承する必要があります。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").inherits_from("ABC")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'
--fcc-editable-region--

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
