---
id: 68e5293bd00d2fe134f58983
title: ステップ 11
challengeType: 20
dashedName: step-11
---

# --description--

`apply_discount` クラスに `DiscountStrategy` という名前の2つ目の抽象メソッドを追加してください。`@abstractmethod` デコレーターを使い、パラメータ `self` と `product` を与えてください。

型ヒントを追加してください：`product` は `Product` 型で、メソッドは `float` を返すべきです。メソッド本体では `pass` を使ってください。

# --hints--

`DiscountStrategy` クラスには `apply_discount` という名前のメソッドがあるべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").has_function("apply_discount")`)
})
```

`apply_discount` メソッドには `@abstractmethod` デコレーターが付いているべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_decorators("abstractmethod")`)
})
```

`apply_discount` メソッドにはパラメータ `self` と `product: Product` があるべきです。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_args("self, product:Product")`)
})
```

`apply_discount` メソッドには `-> float` の戻り型ヒントがあるべきです。

```js
({
  test: () => runPython(`
  assert _Node(_code).find_class("DiscountStrategy").find_function("apply_discount").has_returns("float")
  `)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass
--fcc-editable-region--
    
--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)
```
