---
id: 68e79d91daf8fb4808f93c47
title: ステップ 15
challengeType: 20
dashedName: step-15
---

# --description--

`apply_discount` メソッドを実装したら、割引戦略をテストしてください。`PercentageDiscount` パーセントで `10` のインスタンスを作成し、`discount` という変数に代入します。次に `discount.apply_discount(product)` を使って割引を適用し、結果を出力して割引が正しく動作しているか確認してください。

# --hints--

`discount` のインスタンスを持つ `PercentageDiscount` という名前の変数を作成してください。

```js
({
  test: () => runPython(`
    assert _Node(_code).has_variable("discount")
    assert isinstance(discount, PercentageDiscount)
  `)
})
```

`PercentageDiscount` のインスタンスは `10` パーセントで初期化する必要があります。

```js
({
  test: () => runPython(`assert discount.percent == 10`)
})
```

`discount.apply_discount(product)` を呼び出して結果を出力してください。

```js
({
  test: () => runPython(`assert _Node(_code).has_call("print(discount.apply_discount(product))")`)
})
```

# --seed--

## --seed-contents--

```py
from abc import ABC, abstractmethod

class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price

    def __str__(self) -> str:
        return f'{self.name} - ${self.price}'

class DiscountStrategy(ABC):
    @abstractmethod
    def is_applicable(self, product: Product, user_tier: str) -> bool:
        pass

    @abstractmethod
    def apply_discount(self, product: Product) -> float:
        pass

class PercentageDiscount(DiscountStrategy):
    def __init__(self, percent: int) -> None:
        self.percent = percent

    def is_applicable(self, product: Product, user_tier: str) -> bool:
        return self.percent <= 70
    
    def apply_discount(self, product: Product) -> float:
        return product.price * (1 - self.percent / 100)

--fcc-editable-region--
product = Product('Wireless Mouse', 50.0)
print(product)

--fcc-editable-region--
```
