---
id: 68e9207ddcba053c5cb24413
title: ステップ 2
challengeType: 20
dashedName: step-2
---

# --description--

Pythonでは、型ヒントは任意のもので、変数や関数のデータ型が何であることが期待されているかを他の開発者に伝えるためのものです。ここに、`str`パラメータに`name`の型ヒントを追加した例があります。

```py
def greet(name: str):
    return 'Hello, ' + name
```

`__init__`メソッドのパラメータについて、`name`パラメータに`str`の型ヒントを、`price`パラメータに`float`の型ヒントを追加してください。

# --hints--

`str`クラスの`name`メソッドの`float`パラメータに`price`の型ヒントを、`__init__`パラメータに`Product`の型ヒントを追加してください。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__init__").has_args("self, name: str, price: float")`)
})
```

# --seed--

## --seed-contents--

```py
--fcc-editable-region--
class Product:
    def __init__(self, name, price):
        self.name = name
        self.price = price
--fcc-editable-region--
```
