---
id: 68f119ad905b29ff59d858c8
title: ステップ 5
challengeType: 20
dashedName: step-5
---

# --description--

前のステップと同様に、`__str__` メソッドに、オブジェクトの文字列表現を返すため `str` の戻り値の型ヒントを追加してください。

# --hints--

`__str__` メソッドには `str` の戻り値の型ヒントを付ける必要があります。

```js
({
  test: () => runPython(`assert _Node(_code).find_class("Product").find_function("__str__").has_returns("str")`)
})
```

# --seed--

## --seed-contents--

```py
class Product:
    def __init__(self, name: str, price: float) -> None:
        self.name = name
        self.price = price
--fcc-editable-region--
    def __str__(self):
        return f'{self.name} - ${self.price}'
--fcc-editable-region--
```
