---
id: 61437d575fb98f57fa8f7f36
title: Passaggio 1
challengeType: 0
dashedName: step-1
demoType: onLoad
---

# --description--

Inizia con il tuo codice boilerplate HTML standard. Aggiungi una dichiarazione `DOCTYPE`, un elemento `html` che specifica che questa pagina è in inglese, un elemento `head` e un elemento `body`.

Aggiungi un tag `<meta>` con il `charset` appropriato e un tag `<meta>` per la responsività mobile all'interno dell'elemento `head`.

# --hints--

Il tuo codice dovrebbe contenere il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE/gi));
```

Dovresti includere uno spazio dopo il riferimento `DOCTYPE`.

```js
assert(code.match(/<!DOCTYPE\s+/gi));
```

Dovresti definire il tipo di documento come `html`.

```js
assert(code.match(/<!DOCTYPE\s+html/gi));
```

Dovresti chiudere la dichiarazione `DOCTYPE` con un `>` dopo il tipo.

```js
assert(code.match(/<!DOCTYPE\s+html\s*>/gi));
```

Dovresti avere un tag di apertura `<html>`, e dovrebbe avere `lang` di `en`.

```js
assert(code.match(/<html\s+lang\s*=\s*('|")en\1\s*>/gi));
```

Il tuo elemento `html` dovrebbe avere un tag di chiusura.

```js
assert(code.match(/<\/html\s*>/));
```

La tua dichiarazione `DOCTYPE` dovrebbe essere all'inizio del tuo HTML.

```js
assert(__helpers.removeHtmlComments(code).match(/^\s*<!DOCTYPE\s+html\s*>/i));
```

Dovresti avere un tag di apertura `<head>`.

```js
assert(code.match(/<head\s*>/i));
```

Dovresti avere un tag di chiusura `</head>`.

```js
assert(code.match(/<\/head\s*>/i));
```

Dovresti avere un tag di apertura `<body>`.

```js
assert(code.match(/<body\s*>/i));
```

Dovresti avere un tag di chiusura `</body>`.

```js
assert(code.match(/<\/body\s*>/i));
```

Gli elementi `head` e `body` dovrebbero essere fratelli.

```js
assert(document.querySelector('head')?.nextElementSibling?.localName === 'body');
```

L'elemento `head` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'head'));
```

L'elemento `body` dovrebbe essere all'interno dell'elemento `html`.

```js
assert([...document.querySelector('html')?.children].some(x => x?.localName === 'body'));
```

Dovresti avere due elementi `meta`.

```js
const meta = document.querySelectorAll('meta');
assert(meta?.length === 2);
```

Uno degli elementi `meta` dovrebbe avere un `name` impostato su `viewport`, e `content` impostato su `width=device-width, initial-scale=1.0`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => m?.getAttribute('name') === 'viewport' && m?.getAttribute('content') === 'width=device-width, initial-scale=1.0' && !m?.getAttribute('charset'));
assert.exists(target);
```

L'altro elemento `meta` dovrebbe avere l'attributo `charset` impostato su `UTF-8`.

```js
const meta = [...document.querySelectorAll('meta')];
const target = meta?.find(m => !m?.getAttribute('name') && !m?.getAttribute('content') && m?.getAttribute('charset')?.toLowerCase() === 'utf-8');
assert.exists(target);
```

# --seed--

## --seed-contents--

```html
--fcc-editable-region--

--fcc-editable-region--
```

```css

```
