---
id: 6143862a5e54455d262c212e
title: Passaggio 3
challengeType: 0
dashedName: step-3
---

# --description--

All'interno del tuo `body`, crea un elemento `main`. Poi in quell'elemento, crea un `section` con un insieme `class` impostato su `heading`.

# --hints--

Dovresti avere un elemento `main`.

```js
assert.exists(document.querySelector('main'));
```

Il tuo elemento `main` dovrebbe essere all'interno del tuo elemento `body`.

```js
assert(document.querySelector('main')?.parentElement?.localName === 'body');
```

Dovresti avere un elemento `section`.

```js
assert.exists(document.querySelector('section'));
```

Il tuo elemento `section` dovrebbe essere all'interno del tuo elemento `main`.

```js
assert(document.querySelector('section')?.parentElement?.localName === 'main');
```

Il tuo elemento `section` dovrebbe avere la proprietà `class` impostata su `heading`.

```js
assert(document.querySelector('section')?.classList?.contains('heading'));
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--

--fcc-editable-region--
  </body>
</html>
```

```css

```
