---
id: 614387cbefeeba5f3654a291
title: Passaggio 7
challengeType: 0
dashedName: step-7
---

# --description--

Dopo il tuo elemento `header`, crea un `div` con la proprietà `class` impostata su `author`.

All’interno di quel `div`, crea un elemento `p` con la proprietà `class` impostata su `author-name` e dagli il testo `By freeCodeCamp`. Avvolgi la parte `freeCodeCamp` in un elemento `a` con la proprietà `href` impostata su `https://freecodecamp.org` e la proprietà `target` impostata su `_blank`.

Sotto a quello, aggiungi un secondo elemento `p` con la classe `publish-date` e il testo `March 7, 2019`.

# --hints--

Dovresti creare un nuovo elemento `div`.

```js
assert.exists(document.querySelector('div'));
```

Il tuo elemento `div` dovrebbe venire dopo il tuo elemento `header`.

```js
assert(document.querySelector('div')?.previousElementSibling?.localName === 'header');
```

Il tuo elemento `div` dovrebbe avere la proprietà `class` impostata su `author`.

```js
assert(document.querySelector('div')?.className === 'author');
```

Dovresti creare due nuovi elementi `p`.

```js
assert(document.querySelectorAll('p')?.length === 3);
```

I tuoi due nuovi elementi `p` dovrebbero essere all’interno del tuo elemento `div`.

```js
assert.exists(document.querySelector('div')?.querySelectorAll('p')?.length === 2);
```

Il tuo primo nuovo elemento `p` dovrebbe avere una proprietà `class` impostata su `author-name`.

```js
assert(document.querySelector('div')?.querySelector('p')?.className === 'author-name');
```

Il tuo primo nuovo elemento `p` dovrebbe avere il testo `By freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.innerText === 'By freeCodeCamp');
```

Il tuo secondo nuovo elemento `p` dovrebbe avere una proprietà `class` impostata su `publish-date`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.className === 'publish-date');
```

Il tuo secondo nuovo elemento `p` dovrebbe avere il testo `March 7, 2019`.

```js
assert(document.querySelector('div')?.querySelectorAll('p')?.[1]?.innerText === 'March 7, 2019');
```

Dovresti creare un nuovo elemento `a`.

```js
assert.exists(document.querySelector('a'));
```

Il tuo elemento `a` dovrebbe essere all’interno del tuo primo nuovo elemento `p`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.localName === 'a');
```

Il tuo elemento `a` dovrebbe avere la proprietà `href` impostata su `https://freecodecamp.org`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('href') === 'https://freecodecamp.org');
```

Il tuo elemento `a` dovrebbe avere la proprietà `target` impostata su `_blank`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.getAttribute('target') === '_blank');
```

Il tuo elemento `a` dovrebbe circondare il testo `freeCodeCamp`.

```js
assert(document.querySelector('div')?.querySelector('p')?.firstElementChild?.innerText === 'freeCodeCamp');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
            width="400"
          />
          <h1 class="hero-title">OUR NEW CURRICULUM</h1>
          <p class="hero-subtitle">
            Our efforts to restructure our curriculum with a more project-based
            focus
          </p>
        </header>
--fcc-editable-region--

      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
