---
id: 6165d3b702a5d92ad970b30c
title: Passaggio 5
challengeType: 0
dashedName: step-5
---

# --description--

Dopo il tuo elemento `img`, aggiungi un elemento `h1` con il set `class` impostato su `hero-title` e il testo impostato su `OUR NEW CURRICULUM`, seguito da un elemento `p` con il set `class` impostato su `hero-subtitle` e il testo impostato su `Our efforts to restructure our curriculum with a more project-based focus`.

# --hints--

Dovresti creare un elemento `h1`.

```js
assert.exists(document.querySelector('h1'));
```

Il tuo elemento `h1` dovrebbe venire dopo il tuo elemento `img`.

```js
assert(document.querySelector('h1')?.previousElementSibling?.localName === 'img');
```

Il tuo elemento `h1` dovrebbe avere il set `class` impostato su `hero-title`.

```js
assert(document.querySelector('h1')?.className === 'hero-title');
```

Il tuo elemento `h1` dovrebbe avere il testo impostato su `OUR NEW CURRICULUM`.

```js
assert(document.querySelector('h1')?.textContent === 'OUR NEW CURRICULUM');
```

Dovresti creare un nuovo elemento `p`.

```js
assert.exists(document.querySelector('p'));
```

Il tuo elemento `p` dovrebbe venire dopo il tuo elemento `h1`.

```js
assert(document.querySelector('p')?.previousElementSibling?.localName === 'h1');
```

Il tuo elemento `p` dovrebbe avere il set `class` impostato su `hero-subtitle`.

```js
assert(document.querySelector('p')?.className === 'hero-subtitle');
```

Il tuo elemento `p` dovrebbe avere il testo impostato su `Our efforts to restructure our curriculum with a more project-based focus`.

```js
assert.equal(document.querySelector('p')?.textContent?.trim()?.replace(/\s{2,}/, ' '), 'Our efforts to restructure our curriculum with a more project-based focus');
```

# --seed--

## --seed-contents--

```html
<!DOCTYPE html>
<html lang="en">
  <head>
    <meta charset="UTF-8" />
    <meta name="viewport" content="width=device-width, initial-scale=1.0" />
    <title>Magazine</title>
    <link
      href="https://fonts.googleapis.com/css?family=Anton%7CBaskervville%7CRaleway&display=swap"
      rel="stylesheet"
    />
    <link
      rel="stylesheet"
      href="https://use.fontawesome.com/releases/v5.8.2/css/all.css"
    />
    <link rel="stylesheet" href="styles.css" />
  </head>
  <body>
--fcc-editable-region--
    <main>
      <section class="heading">
        <header class="hero">
          <img
            src="https://cdn.freecodecamp.org/platform/universal/fcc_meta_1920X1080-indigo.png"
            alt="freecodecamp logo"
            loading="lazy"
            class="hero-img"
          />
        </header>
      </section>
    </main>
--fcc-editable-region--
  </body>
</html>
```

```css

```
