---
id: 594810f028c0303b75339acc
title: مشكلة ABC
challengeType: 1
forumTopicId: 302220
dashedName: abc-problem
---

# --description--

لديك مجموعة من كتل ABC (مثل كتل الأبجدية في الطفولة). هناك 20 كتلة، وكل كتلة تحمل حرفين. الأبجدية كاملة مضمونة بين جميع جوانب الكتل. مجموعة الكتل النموذجية:

<pre>(B O)
(X K)
(D Q)
(C P)
(N A)
(G T)
(R E)
(T G)
(Q D)
(F S)
(J W)
(H U)
(V I)
(A N)
(O B)
(E R)
(F S)
(L Y)
(P C)
(Z M)
</pre>

# --instructions--

نفذ دالة تأخذ سلسلة نصية (كلمة) وتحدد ما إذا كان يمكن تهجئة الكلمة باستخدام مجموعة الكتل المعطاة.

بعض القواعد التي يجب وضعها في الاعتبار:

<ul>
<li>بمجرد استخدام حرف على كتلة، لا يمكن استخدام تلك الكتلة مرة أخرى.</li>
  <li>يجب أن تكون الدالة غير حساسة لحالة الأحرف.</li>
</ul>

# --before-each--

```js
const words = ['bark', 'BooK', 'TReAT', 'COMMON', 'squAD', 'conFUSE'];
```

# --hints--

يجب أن تكون `canMakeWord` دالة.

```js
assert(typeof canMakeWord === 'function');
```

يجب أن تُرجع `canMakeWord` قيمة منطقية.

```js
assert(typeof canMakeWord('hi') === 'boolean');
```

يجب أن تُرجع `canMakeWord("bark")` القيمة true.

```js
assert(canMakeWord(words[0]));
```

يجب أن تُرجع `canMakeWord("BooK")` القيمة false.

```js
assert(!canMakeWord(words[1]));
```

يجب أن تُرجع `canMakeWord("TReAT")` القيمة true.

```js
assert(canMakeWord(words[2]));
```

يجب أن تُرجع `canMakeWord("COMMON")` القيمة false.

```js
assert(!canMakeWord(words[3]));
```

يجب أن تُرجع `canMakeWord("squAD")` القيمة true.

```js
assert(canMakeWord(words[4]));
```

يجب أن تُرجع `canMakeWord("conFUSE")` القيمة true.

```js
assert(canMakeWord(words[5]));
```

# --seed--

## --seed-contents--

```js
function canMakeWord(word) {

}
```

# --solutions--

```js
function canMakeWord(word) {
  const characters = 'BO XK DQ CP NA GT RE TG QD FS JW HU VI AN OB ER FS LY PC ZM';
  const blocks = characters.split(' ').map(pair => pair.split(''));

  const letters = [...word.toUpperCase()];
  let length = letters.length;
  const copy = new Set(blocks);

  letters.forEach(letter => {
    for (let block of copy) {
      const index = block.indexOf(letter);

      if (index !== -1) {
        length--;
        copy.delete(block);
        break;
      }
    }
  });
  return !length;
}
```
