---
id: 594810f028c0303b75339ad8
title: مصفوفة متعرجة
challengeType: 1
forumTopicId: 302348
dashedName: zig-zag-matrix
---

# --description--

المصفوفة "المتعرجة" هي ترتيب مربع لأول $N^2$ أعداد صحيحة، حيث تزداد الأعداد بالتتابع أثناء التنقل بشكل متعرج على الأقطار المضادة للمصفوفة.

على سبيل المثال، للإدخال `5`، يجب إنتاج النتيجة التالية:

<pre>
 0  1  5  6 14
 2  4  7 13 15
 3  8 12 16 21
 9 11 17 20 22
10 18 19 23 24
</pre>

# --instructions--

اكتب دالة تأخذ حجم المصفوفة المتعرجة، وتُرجع المصفوفة المقابلة كمصفوفة ثنائية الأبعاد.

# --before-each--

```js
const zm1 = [[0]];
const zm2 = [[0, 1], [2, 3]];
const zm5 = [
  [0, 1, 5, 6, 14],
  [2, 4, 7, 13, 15],
  [3, 8, 12, 16, 21],
  [9, 11, 17, 20, 22],
  [10, 18, 19, 23, 24]
];
```

# --hints--

يجب أن تكون ZigZagMatrix دالة.

```js
assert.equal(typeof ZigZagMatrix, 'function');
```

يجب أن تُرجع ZigZagMatrix مصفوفة.

```js
assert.equal(typeof ZigZagMatrix(1), 'object');
```

يجب أن تُرجع ZigZagMatrix مصفوفة من مصفوفات مسكنة.

```js
assert.equal(typeof ZigZagMatrix(1)[0], 'object');
```

يجب أن تُرجع ZigZagMatrix(1) \[[0]].

```js
assert.deepEqual(ZigZagMatrix(1), zm1);
```

يجب أن تُرجع ZigZagMatrix(2) \[[0, 1], [2, 3]].

```js
assert.deepEqual(ZigZagMatrix(2), zm2);
```

يجب أن تُرجع ZigZagMatrix(5) المصفوفة المحددة.

```js
assert.deepEqual(ZigZagMatrix(5), zm5);
```

# --seed--

## --seed-contents--

```js
function ZigZagMatrix(n) {

  return [[], []];
}
```

# --solutions--

```js
function ZigZagMatrix(n) {
  const mtx = [];
  for (let i = 0; i < n; i++) {
    mtx[i] = [];
  }

  let i = 1;
  let j = 1;
  for (let e = 0; e < n * n; e++) {
    mtx[i - 1][j - 1] = e;
    if ((i + j) % 2 === 0) {
      // Even stripes
      if (j < n) j++;
      else i += 2;
      if (i > 1) i--;
    } else {
      // Odd stripes
      if (i < n) i++;
      else j += 2;
      if (j > 1) j--;
    }
  }
  return mtx;
}
```
