---
id: 594810f028c0303b75339ad1
title: الأعداد السعيدة
challengeType: 1
forumTopicId: 302280
dashedName: happy-numbers
---

# --description--

يُعرّف العدد السعيد بالعملية التالية:

ابدأ بأي عدد صحيح موجب، واستبدل العدد بمجموع مربعات أرقامه، وكرر العملية حتى يصبح العدد `1` (حيث يبقى)، أو يدخل في دورة لا نهائية لا تشمل `1`. الأعداد التي تنتهي هذه العملية عندها بـ `1` هي أعداد سعيدة، أما التي لا تنتهي بـ `1` فهي أعداد غير سعيدة.

# --instructions--

نفذ دالة تُرجع true إذا كان العدد سعيدًا، أو false إذا لم يكن كذلك.

# --hints--

يجب أن تكون `happy` دالة.

```js
assert(typeof happy === 'function');
```

يجب أن تُرجع `happy(1)` قيمة منطقية.

```js
assert(typeof happy(1) === 'boolean');
```

يجب أن تُرجع `happy(1)` القيمة `true`.

```js
assert(happy(1));
```

يجب أن تُرجع `happy(2)` القيمة `false`.

```js
assert(!happy(2));
```

يجب أن تُرجع `happy(7)` القيمة `true`.

```js
assert(happy(7));
```

يجب أن تُرجع `happy(10)` القيمة `true`.

```js
assert(happy(10));
```

يجب أن تُرجع `happy(13)` القيمة `true`.

```js
assert(happy(13));
```

يجب أن تُرجع `happy(19)` القيمة `true`.

```js
assert(happy(19));
```

يجب أن تُرجع `happy(23)` القيمة `true`.

```js
assert(happy(23));
```

يجب أن تُرجع `happy(28)` القيمة `true`.

```js
assert(happy(28));
```

يجب أن تُرجع `happy(31)` القيمة `true`.

```js
assert(happy(31));
```

يجب أن تُرجع `happy(32)` القيمة `true`.

```js
assert(happy(32));
```

يجب أن تُرجع `happy(33)` القيمة `false`.

```js
assert(!happy(33));
```

# --seed--

## --seed-contents--

```js
function happy(number) {

}
```

# --solutions--

```js
function happy (number) {
  let m;
  let digit;
  const cycle = [];

  while (number !== 1 && cycle[number] !== true) {
    cycle[number] = true;
    m = 0;
    while (number > 0) {
      digit = number % 10;
      m += Math.pow(digit, 2);
      number = (number - digit) / 10;
    }
    number = m;
  }
  return (number === 1);
}
```
