---
id: 594810f028c0303b75339acb
title: 100 باب
challengeType: 1
forumTopicId: 302217
dashedName: 100-doors
---

# --description--

هناك 100 باب مصطفة في صف وكلها مغلقة في البداية. تقوم بتمرير 100 مرة على الأبواب. في المرور الأول، زر كل باب و"بدّل" حالته (إذا كان الباب مغلقًا، افتحه؛ وإذا كان مفتوحًا، أغلقه). في المرور الثاني، زر كل باب ثاني فقط (أي الباب رقم 2، 4، 6، ...) وبدّل حالته. في المرور الثالث، زر كل باب ثالث (أي الباب رقم 3، 6، 9، ...) وهكذا، حتى تزور الباب رقم 100 فقط.

# --instructions--

نفذ دالة لتحديد حالة الأبواب بعد المرور الأخير. أرجع النتيجة النهائية في مصفوفة، بحيث يتضمن الرقم فقط للأبواب المفتوحة.

# --hints--

`getFinalOpenedDoors` يجب أن تكون دالة.

```js
assert(typeof getFinalOpenedDoors === 'function');
```

`getFinalOpenedDoors` يجب أن تُرجع مصفوفة.

```js
assert(Array.isArray(getFinalOpenedDoors(100)));
```

`getFinalOpenedDoors` يجب أن تُنتج النتيجة الصحيحة.

```js
const solution = [1, 4, 9, 16, 25, 36, 49, 64, 81, 100];
assert.deepEqual(getFinalOpenedDoors(100), solution);
```

# --seed--

## --seed-contents--

```js
function getFinalOpenedDoors(numDoors) {

}
```

# --solutions--

```js
function getFinalOpenedDoors(numDoors) {
  // this is the final pattern (always squares).
  // thus, the most efficient solution simply returns an array of squares up to numDoors).
  const finalState = [];
  let i = 1;
  while (Math.pow(i, 2) <= numDoors) {
    finalState.push(Math.pow(i, 2));
    i++;
  }
  return finalState;
}
```
