---
id: 594810f028c0303b75339ace
title: مصنع المجمع
challengeType: 1
forumTopicId: 302222
dashedName: accumulator-factory
---

# --description--

مشكلة طرحها بول غراهام هي إنشاء دالة تأخذ معلمة رقمية واحدة وتُرجع دالة أخرى تكون مجمعًا. دالة المجمع المُرجعة تأخذ بدورها معلمة رقمية واحدة، وتُرجع مجموع كل القيم الرقمية التي تم تمريرها حتى الآن إلى ذلك المجمع (بما في ذلك القيمة الأولية التي تم تمريرها عند إنشاء المجمع).

# --instructions--

أنشئ دالة تأخذ رقمًا $n$ وتُولد دوال مجمع تُرجع مجموع كل الأرقام التي تم تمريرها إليها على الإطلاق.

**القواعد:**

لا تستخدم متغيرات عامة.

**تلميح:**

الإغلاقات تحفظ الحالة الخارجية.

# --hints--

يجب أن يكون `accumulator` دالة.

```js
assert(typeof accumulator === 'function');
```

يجب أن تُرجع `accumulator(0)` دالة.

```js
assert(typeof accumulator(0) === 'function');
```

يجب أن تُرجع `accumulator(0)(2)` رقمًا.

```js
assert(typeof accumulator(0)(2) === 'number');
```

تمرير القيم 3، -4، 1.5، و5 يجب أن يُرجع 5.5.

```js
assert(testFn && testFn(5) === 5.5);
```

# --before-each--

```js
const testFn = typeof accumulator === 'function' && accumulator(3);
if (testFn) {
  testFn(-4);
  testFn(1.5);
}
```

# --seed--

## --seed-contents--

```js
function accumulator(sum) {

}
```

# --solutions--

```js
function accumulator(sum) {
  return function(n) {
    return sum += n;
  };
}
```
