---
id: 594810f028c0303b75339acd
title: "تصنيفات الأعداد الوفيرة والناقصة والكاملة"
challengeType: 1
forumTopicId: 302221
dashedName: abundant-deficient-and-perfect-number-classifications
---

# --description--

تعرف هذه التصنيفات ثلاثة أنواع من الأعداد الصحيحة الموجبة بناءً على القواسم الصحيحة الخاصة بها.

لتكن $P(n)$ مجموع القواسم الصحيحة الخاصة للعدد `n` حيث القواسم الصحيحة الخاصة هي كل الأعداد الصحيحة الموجبة لـ `n` ما عدا `n` نفسه.

إذا كان `P(n) < n` فإن `n` يصنف كـ `deficient`

إذا كان `P(n) === n` فإن `n` يصنف كـ `perfect`

إذا كان `P(n) > n` فإن `n` يصنف كـ `abundant`

**مثال**: العدد `6` له قواسم صحيحة خاصة هي `1` و `2` و `3`. حيث أن `1 + 2 + 3 = 6`، إذًا `6` يصنف كـ العدد الكامل.

# --instructions--

نفذ دالة تحسب عدد الأعداد الصحيحة من `1` إلى `num` (شاملة) التي تنتمي لكل من التصنيفات الثلاثة. أخرج النتيجة كمصفوفة بالتنسيق التالي `[deficient, perfect, abundant]`.

# --hints--

يجب أن تكون `getDPA` دالة.

```js
assert(typeof getDPA === 'function');
```

يجب أن تعيد `getDPA(5000)` مصفوفة.

```js
assert(Array.isArray(getDPA(5000)));
```

يجب أن يكون طول المصفوفة التي تعيدها `getDPA(5000)` هو `3`.

```js
assert(getDPA(5000).length === 3);
```

يجب أن تعيد `getDPA(5000)` `[3758, 3, 1239]`.

```js
assert.deepEqual(getDPA(5000), [3758, 3, 1239]);
```

يجب أن تعيد `getDPA(10000)` `[7508, 4, 2488]`.

```js
assert.deepEqual(getDPA(10000), [7508, 4, 2488]);
```

يجب أن تعيد `getDPA(20000)` `[15043, 4, 4953]`.

```js
assert.deepEqual(getDPA(20000), [15043, 4, 4953]);
```

# --seed--

## --seed-contents--

```js
function getDPA(num) {

}
```

# --solutions--

```js
function getDPA(num) {
  const dpa = [1, 0, 0];
  for (let n = 2; n <= num; n += 1) {
    let ds = 1;
    const e = Math.sqrt(n);
    for (let d = 2; d < e; d += 1) {
      if (n % d === 0) {
        ds += d + (n / d);
      }
    }
    if (n % e === 0) {
      ds += e;
    }
    dpa[ds < n ? 0 : ds === n ? 1 : 2] += 1;
  }
  return dpa;
}
```
