---
id: 594810f028c0303b75339ad3
title: الضرب الداخلي للمتجه
challengeType: 1
forumTopicId: 302343
dashedName: vector-dot-product
---

# --description--

يمكن أن يحتوي المتجه على قيمة واحدة أو أكثر يتم تمثيلها بواسطة مجموعة مرتبة. قد تكون الأمثلة (x)، أو (x, y)، أو (x, y, z).

# --instructions--

اكتب دالة تأخذ متجهين (ممثّلين كمصفوفات أحادية البعد) كمُدخَل وتحسب الضرب الداخلي لهما. يجب أن تُرجع الدالة `null` عند وجود مُدخلات غير صالحة مثل المتجهات ذات الأطوال المختلفة أو تمرير أي شيء غير متجهين.

# --hints--

يجب أن تكون `dotProduct` دالة.

```js
assert.equal(typeof dotProduct, 'function');
```

يجب أن تُرجع `dotProduct()` القيمة `null`.

```js
assert.equal(dotProduct(), null);
```

يجب أن تُرجع `dotProduct([1], [1])` القيمة `1`.

```js
assert.equal(dotProduct([1], [1]), 1);
```

يجب أن تُرجع `dotProduct([1], [1, 2])` القيمة `null`.

```js
assert.equal(dotProduct([1], [1, 2]), null);
```

يجب أن تُرجع `dotProduct([1, 3, -5], [4, -2, -1])` القيمة `3`.

```js
assert.equal(dotProduct([1, 3, -5], [4, -2, -1]), 3);
```

يجب أن تُرجع `dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1])` القيمة `null`.

```js
assert.equal(dotProduct([3, 2, 1], [2, 4, 2], [5, 3, 1]), null);
```

يجب أن تُرجع `dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ])` القيمة `360`.

```js
assert.equal(dotProduct([ 0, 3, 6, 9, 12 ], [ 0, 4, 8, 12, 16 ]), 360);
```

# --seed--

## --seed-contents--

```js
function dotProduct(...vectors) {

}
```

# --solutions--

```js
function dotProduct(...vectors) {
  if (!vectors || !vectors.length || vectors.length > 2 || vectors[0].length !== vectors[1].length) {
    return null;
  }
  const vectorLen = vectors[0].length;

  let prod = 0;
  let sum = 0;
  let j = vectorLen;
  let i = 2;
  // Sum terms
  while (j--) {
    i = 2;
    prod = 1;

    while (i--) {
      prod *= vectors[i][j];
    }
    sum += prod;
  }
  return sum;
}
```
