---
id: 594810f028c0303b75339ad5
title: مركب Y
challengeType: 1
forumTopicId: 302345
dashedName: y-combinator
---

# --description--

في <a href="https://www.freecodecamp.org/news/the-principles-of-functional-programming/" target="_blank" rel="noopener noreferrer nofollow">البرمجة الوظيفية</a> الصارمة وحساب اللامبدا، الدوال (تعبيرات اللامبدا) لا تملك حالة ويُسمح لها فقط بالإشارة إلى المعلمات الخاصة بالدوال المحيطة. هذا يستبعد التعريف المعتاد للدالة العودية التي ترتبط بحالة متغير وتُستخدم حالة هذا المتغير داخل جسم الدالة.

مركب Y هو دالة عديمة الحالة بحد ذاته، وعند تطبيقه على دالة عديمة الحالة أخرى، يعيد نسخة عودية من الدالة. مركب Y هو الأبسط ضمن فئة هذه الدوال، التي تُسمى مركبات النقطة الثابتة.

# --instructions--

عرّف دالة مركب Y عديم الحالة واستخدمها لحساب المضروبات. الدالة `factorial(N)` مُعطاة لك بالفعل.

# --hints--

يجب أن يعيد Y دالة.

```js
assert.equal(typeof Y((f) => (n) => n), 'function');
```

يجب أن تُعيد factorial(1) القيمة 1.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(1), 1);
```

يجب أن تُعيد factorial(2) القيمة 2.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(2), 2);
```

يجب أن تُعيد factorial(3) القيمة 6.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(3), 6);
```

يجب أن تُعيد factorial(4) القيمة 24.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(4), 24);
```

يجب أن تُعيد factorial(10) القيمة 3628800.

```js
var factorial = Y(f => n => (n > 1 ? n * f(n - 1) : 1));
assert.equal(factorial(10), 3628800);
```

# --seed--

## --seed-contents--

```js
function Y(f) {
  return function() {

  };
}

var factorial = Y(function(f) {
  return function (n) {
    return n > 1 ? n * f(n - 1) : 1;
  };
});
```

# --solutions--

```js
var Y = f => (x => x(x))(y => f(x => y(y)(x)));
```
