---
id: 594810f028c0303b75339ad7
title: خوارزمية ترقيق Zhang-Suen
challengeType: 1
forumTopicId: 302347
dashedName: zhang-suen-thinning-algorithm
---

# --description--

هذه خوارزمية تُستخدم لترقيق الصور ذات اللونين الأبيض والأسود، أي صورة ببت واحد لكل بكسل. على سبيل المثال، مع صورة مُدخَل مثل:

```js
const testImage1 = [
 '                               ',
 '#########       ########       ',
 '###   ####     ####  ####      ',
 '###    ###     ###    ###      ',
 '###   ####     ###             ',
 '#########      ###             ',
 '### ####       ###    ###      ',
 '###  ####  ### ####  #### ###  ',
 '###   #### ###  ########  ###  ',
 '                               '
];
```

تُنتج الصورة المُرققة التالية:

```js
[ '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               ' ];
```

## الخوارزمية

افترض أن البكسلات السوداء تمثل الواحدات والبكسلات البيضاء تمثل الأصفار، وأن الصورة المُدخلة هي مصفوفة مستطيلة أبعادها N في M من الواحدات والأصفار. تعمل الخوارزمية على جميع البكسلات السوداء P1 التي يمكن أن يكون لها ثمانية جيران. الجيران مرتّبون بالترتيب كما يلي:

$$\begin{array}{|c|c|c|}
  \\hline
  P9 & P2              & P3\\\\ \\hline
  P8 & \boldsymbol{P1} & P4\\\\ \\hline
  P7 & P6              & P5\\\\ \\hline
\end{array}$$

من الواضح أن بكسلات الحدود في الصورة لا يمكن أن يكون لها الجيران الثمانية كاملة.

- عرّف $A(P1)$ = عدد الانتقالات من الأبيض إلى الأسود، ($0 \to 1$) في المتتالية P2, P3, P4, P5, P6, P7, P8, P9, P2. (لاحظ وجود P2 إضافي في النهاية - المتتالية دائرية).
- عرّف $B(P1)$ = عدد جيران البكسل السوداء لـ P1. ($= \\sum(P2 \ldots P9)$)

**الخطوة 1:**

تُختبر جميع البكسلات وتُلاحظ فقط البكسلات التي تحقق جميع الشروط التالية (في نفس الوقت) في هذه المرحلة.

1. البكسل أسود وله ثمانية جيران
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. على الأقل واحد من $P2$، $P4$ و $P6$ أبيض
5. على الأقل واحد من $P4$، $P6$ و $P8$ أبيض

بعد المرور على الصورة وجمع كل البكسلات التي تحقق جميع شروط الخطوة 1، تُحول كل هذه البكسلات إلى اللون الأبيض.

**الخطوة 2:**

تُختبر جميع البكسلات مرة أخرى وتُلاحظ فقط البكسلات التي تحقق جميع الشروط التالية في هذه المرحلة.

1. البكسل أسود وله ثمانية جيران
2. $2 \le B(P1) \le 6$
3. $A(P1) = 1$
4. على الأقل واحد من $P2$، $P4$ و $P8$ أبيض
5. على الأقل واحد من $P2$، $P6$ و $P8$ أبيض

بعد المرور على الصورة وجمع كل البكسلات التي تحقق جميع شروط الخطوة 2، تُحول كل هذه البكسلات إلى اللون الأبيض مرة أخرى.

**التكرار:**

إذا تم تغيير أي بكسلات في هذه الجولة من الخطوة 1 أو الخطوة 2، تُكرر كل الخطوات حتى لا يتغير أي بكسل في الصورة.

# --instructions--

اكتب روتينًا لتنفيذ ترقيق Zhang-Suen على `image` المُعطاة، وهي مصفوفة من السلاسل النصية، حيث تمثل كل سلسلة سطرًا واحدًا من الصورة. في السلسلة، يمثل `#` البكسل الأسود، والمسافة تمثل البكسل الأبيض. يجب أن تُرجع الدالة الصورة المرققة باستخدام نفس التمثيل.

# --before-each--

```js
const _testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
const expected1 = [
  '                               ',
  '########         ######        ',
  '#      #        ##             ',
  '#       #       #              ',
  '#      #        #              ',
  '###### #        #              ',
  '#     ##        #              ',
  '#      #    #   ##    ##   #   ',
  '#       #         ####         ',
  '                               '
];
const _testImage2 = [
  '                                                          ',
  ' #################                   #############        ',
  ' ##################               ################        ',
  ' ###################            ##################        ',
  ' ########     #######          ###################        ',
  '   ######     #######         #######       ######        ',
  '   ######     #######        #######                      ',
  '   #################         #######                      ',
  '   ################          #######                      ',
  '   #################         #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######        #######                      ',
  '   ######     #######         #######       ######        ',
  ' ########     #######          ###################        ',
  ' ########     ####### ######    ################## ###### ',
  ' ########     ####### ######      ################ ###### ',
  ' ########     ####### ######         ############# ###### ',
  '                                                          '
];
const expected2 = [
  '                                                          ',
  '                                                          ',
  '    # ##########                       #######            ',
  '     ##        #                   ####       #           ',
  '     #          #                 ##                      ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     ############               #                         ',
  '     #          #               #                         ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #          #                #                        ',
  '     #                            ##                      ',
  '     #                             ############           ',
  '                       ###                          ###   ',
  '                                                          ',
  '                                                          '
];
```

# --hints--

`thinImage` يجب أن تكون دالة.

```js
assert.equal(typeof thinImage, 'function');
```

`thinImage` يجب أن تُرجع مصفوفة.

```js
assert(Array.isArray(thinImage(_testImage1)));
```

`thinImage` يجب أن تُرجع مصفوفة من السلاسل النصية.

```js
assert.equal(typeof thinImage(_testImage1)[0], 'string');
```

`thinImage(testImage1)` يجب أن تُرجع صورة مرققة كما في المثال.

```js
assert.deepEqual(thinImage(_testImage1), expected1);
```

`thinImage(testImage2)` يجب أن تُرجع صورة مرققة.

```js
assert.deepEqual(thinImage(_testImage2), expected2);
```

# --seed--

## --seed-contents--

```js
function thinImage(image) {

}

const testImage1 = [
  '                               ',
  '#########       ########       ',
  '###   ####     ####  ####      ',
  '###    ###     ###    ###      ',
  '###   ####     ###             ',
  '#########      ###             ',
  '### ####       ###    ###      ',
  '###  ####  ### ####  #### ###  ',
  '###   #### ###  ########  ###  ',
  '                               '
];
```

# --solutions--

```js
function Point(x, y) {
  this.x = x;
  this.y = y;
}

const ZhangSuen = (function () {
  function ZhangSuen() {
  }

  ZhangSuen.nbrs = [[0, -1], [1, -1], [1, 0], [1, 1], [0, 1], [-1, 1], [-1, 0], [-1, -1], [0, -1]];

  ZhangSuen.nbrGroups = [[[0, 2, 4], [2, 4, 6]], [[0, 2, 6], [0, 4, 6]]];

  ZhangSuen.toWhite = [];

  ZhangSuen.main = function (image) {
    ZhangSuen.grid = new Array(image);
    for (let r = 0; r < image.length; r++) {
      ZhangSuen.grid[r] = image[r].split('');
    }
    ZhangSuen.thinImage();
    return ZhangSuen.getResult();
  };

  ZhangSuen.thinImage = function () {
    let firstStep = false;
    let hasChanged;
    do {
      hasChanged = false;
      firstStep = !firstStep;
      for (let r = 1; r < ZhangSuen.grid.length - 1; r++) {
        for (let c = 1; c < ZhangSuen.grid[0].length - 1; c++) {
          if (ZhangSuen.grid[r][c] !== '#') {
            continue;
          }
          const nn = ZhangSuen.numNeighbors(r, c);
          if (nn < 2 || nn > 6) {
            continue;
          }
          if (ZhangSuen.numTransitions(r, c) !== 1) {
            continue;
          }
          if (!ZhangSuen.atLeastOneIsWhite(r, c, firstStep ? 0 : 1)) {
            continue;
          }
          ZhangSuen.toWhite.push(new Point(c, r));
          hasChanged = true;
        }
      }
      for (let i = 0; i < ZhangSuen.toWhite.length; i++) {
        const p = ZhangSuen.toWhite[i];
        ZhangSuen.grid[p.y][p.x] = ' ';
      }
      ZhangSuen.toWhite = [];
    } while ((firstStep || hasChanged));
  };

  ZhangSuen.numNeighbors = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === '#') {
        count++;
      }
    }
    return count;
  };

  ZhangSuen.numTransitions = function (r, c) {
    let count = 0;
    for (let i = 0; i < ZhangSuen.nbrs.length - 1; i++) {
      if (ZhangSuen.grid[r + ZhangSuen.nbrs[i][1]][c + ZhangSuen.nbrs[i][0]] === ' ') {
        if (ZhangSuen.grid[r + ZhangSuen.nbrs[i + 1][1]][c + ZhangSuen.nbrs[i + 1][0]] === '#') {
          count++;
        }
      }
    }
    return count;
  };

  ZhangSuen.atLeastOneIsWhite = function (r, c, step) {
    let count = 0;
    const group = ZhangSuen.nbrGroups[step];
    for (let i = 0; i < 2; i++) {
      for (let j = 0; j < group[i].length; j++) {
        const nbr = ZhangSuen.nbrs[group[i][j]];
        if (ZhangSuen.grid[r + nbr[1]][c + nbr[0]] === ' ') {
          count++;
          break;
        }
      }
    }
    return count > 1;
  };

  ZhangSuen.getResult = function () {
    const result = [];
    for (let i = 0; i < ZhangSuen.grid.length; i++) {
      const row = ZhangSuen.grid[i].join('');
      result.push(row);
    }
    return result;
  };
  return ZhangSuen;
}());

function thinImage(image) {
  return ZhangSuen.main(image);
}
```
