---
id: 594810f028c0303b75339ad6
title: تمثيل رقم زيكيندورف
challengeType: 1
forumTopicId: 302346
dashedName: zeckendorf-number-representation
---

# --description--

تمامًا كما يمكن تمثيل الأعداد في تدوين موضعي كمجاميع مضاعفات قوى العشرة (عشري) أو الاثنين (ثنائي)؛ يمكن تمثيل كل الأعداد الصحيحة الموجبة كمجموع واحد أو صفر مضروبًا في الأعضاء المميزة من متتالية فيبوناتشي. تذكر أن أول ستة أعداد فيبوناتشي مميزة هي: `1, 2, 3, 5, 8, 13`.

يمكن كتابة العدد العشري أحد عشر كـ `0*13 + 1*8 + 0*5 + 1*3 + 0*2 + 0*1` أو `010100` في التدوين الموضعي حيث تمثل الأعمدة الضرب في عضو معين من المتتالية. تُحذف الأصفار البادئة ليصبح 11 عشريًا `10100`. لكن 10100 ليست الطريقة الوحيدة لتكوين 11 من أعداد فيبوناتشي، إذ إن `0*13 + 1*8 + 0*5 + 0*3 + 1*2 + 1*1` أو 010011 تمثل أيضًا العدد 11 عشريًا. للرقم الحقيقي لزيكيندورف هناك قيد إضافي وهو *عدم استخدام عددين متتاليين من فيبوناتشي* مما يؤدي إلى الحل الفريد السابق.

# --instructions--

اكتب دالة تولد وتُرجع تمثيل رقم زيكيندورف للعدد `n`.

# --hints--

`zeckendorf` يجب أن تكون دالة.

```js
assert.equal(typeof zeckendorf, 'function');
```

`zeckendorf(0)` يجب أن تُرجع `0`.

```js
assert.equal(zeckendorf(0), 0);

```

`zeckendorf(1)` يجب أن تُرجع `1`.

```js
assert.equal(zeckendorf(1), 1);
```

`zeckendorf(2)` يجب أن تُرجع `10`.

```js
assert.equal(zeckendorf(2), 10);
```

`zeckendorf(3)` يجب أن تُرجع `100`.

```js
assert.equal(zeckendorf(3), 100);
```

`zeckendorf(4)` يجب أن تُرجع `101`.

```js
assert.equal(zeckendorf(4), 101);
```

`zeckendorf(5)` يجب أن تُرجع `1000`.

```js
assert.equal(zeckendorf(5), 1000);
```

`zeckendorf(6)` يجب أن تُرجع `1001`.

```js
assert.equal(zeckendorf(6), 1001);
```

`zeckendorf(7)` يجب أن تُرجع `1010`.

```js
assert.equal(zeckendorf(7), 1010);
```

`zeckendorf(8)` يجب أن تُرجع `10000`.

```js
assert.equal(zeckendorf(8), 10000);
```

`zeckendorf(9)` يجب أن تُرجع `10001`.

```js
assert.equal(zeckendorf(9), 10001);
```

`zeckendorf(10)` يجب أن تُرجع `10010`.

```js
assert.equal(zeckendorf(10), 10010);
```

`zeckendorf(11)` يجب أن تُرجع `10100`.

```js
assert.equal(zeckendorf(11), 10100);
```

`zeckendorf(12)` يجب أن تُرجع `10101`.

```js
assert.equal(zeckendorf(12), 10101);
```

`zeckendorf(13)` يجب أن تُرجع `100000`.

```js
assert.equal(zeckendorf(13), 100000);
```

`zeckendorf(14)` يجب أن تُرجع `100001`.

```js
assert.equal(zeckendorf(14), 100001);
```

`zeckendorf(15)` يجب أن تُرجع `100010`.

```js
assert.equal(zeckendorf(15), 100010);
```

`zeckendorf(16)` يجب أن تُرجع `100100`.

```js
assert.equal(zeckendorf(16), 100100);
```

`zeckendorf(17)` يجب أن تُرجع `100101`.

```js
assert.equal(zeckendorf(17), 100101);
```

`zeckendorf(18)` يجب أن تُرجع `101000`.

```js
assert.equal(zeckendorf(18), 101000);
```

`zeckendorf(19)` يجب أن تُرجع `101001`.

```js
assert.equal(zeckendorf(19), 101001);
```

`zeckendorf(20)` يجب أن تُرجع `101010`.

```js
assert.equal(zeckendorf(20), 101010);
```

# --seed--

## --seed-contents--

```js
function zeckendorf(n) {

}
```

# --solutions--

```js
// zeckendorf :: Int -> Int
function zeckendorf(n) {
  const f = (m, x) => (m < x ? [m, 0] : [m - x, 1]);
  return parseInt((n === 0 ? ([0]) :
    mapAccumL(f, n, reverse(
      tail(fibUntil(n))
    ))[1]).join(''));
}

// fibUntil :: Int -> [Int]
let fibUntil = n => {
  const xs = [];
  until(
      ([a]) => a > n,
      ([a, b]) => (xs.push(a), [b, a + b]), [1, 1]
  );
  return xs;
};

let mapAccumL = (f, acc, xs) => (
  xs.reduce((a, x) => {
    const pair = f(a[0], x);

    return [pair[0], a[1].concat(pair[1])];
  }, [acc, []])
);

let until = (p, f, x) => {
  let v = x;
  while (!p(v)) v = f(v);
  return v;
};

const tail = xs => (
   xs.length ? xs.slice(1) : undefined
);

const reverse = xs => xs.slice(0).reverse();
```
