---
id: 594810f028c0303b75339ad2
title: حاصل الضرب الاتجاهي للمتجه
challengeType: 1
forumTopicId: 302342
dashedName: vector-cross-product
---

# --description--

يُعرّف المتجه بأنه ثلاثي الأبعاد ويتم تمثيله بمجموعة مرتبة من ثلاثة أرقام: (X, Y, Z).

# --instructions--

اكتب دالة تأخذ متجهين (مصفوفتين) كمُدخَل وتحسب حاصل ضربهما الاتجاهي. يجب أن تُرجع الدالة `null` عند وجود مُدخلات غير صالحة مثل المتجهات ذات الأطوال المختلفة.

# --hints--

يجب أن تكون `crossProduct` دالة.

```js
assert.equal(typeof crossProduct, 'function');
```

يجب أن تُرجع `crossProduct()` القيمة null.

```js
assert.equal(crossProduct(), null);
```

يجب أن تُرجع `crossProduct([1, 2, 3], [4, 5, 6])` القيمة `[-3, 6, -3]`.

```js
assert.deepEqual(res12, exp12);
```

# --before-each--

```js
const tv1 = [1, 2, 3];
const tv2 = [4, 5, 6];
const res12 = crossProduct(tv1, tv2);
const exp12 = [-3, 6, -3];
```

# --seed--

## --seed-contents--

```js
function crossProduct(a, b) {

}
```

# --solutions--

```js
function crossProduct(a, b) {
  if (!a || !b) {
    return null;
  }

  // Check lengths
  if (a.length !== 3 || b.length !== 3) {
    return null;
  }

  return [
    (a[1] * b[2]) - (a[2] * b[1]),
    (a[2] * b[0]) - (a[0] * b[2]),
    (a[0] * b[1]) - (a[1] * b[0])
  ];
}
```
