---
id: 594810f028c0303b75339acf
title: دالة آكرمان
challengeType: 1
forumTopicId: 302223
dashedName: ackermann-function
---

# --description--

دالة آكرمان مثال كلاسيكي على دالة متكررة، وتتميز خصوصًا لأنها ليست دالة متكررة بدائية. تنمو قيمتها بسرعة كبيرة، وكذلك حجم شجرة الاستدعاءات الخاصة بها.

عادةً ما تُعرف دالة آكرمان كما يلي:

$A(m, n) = \\begin{cases} n+1 & \\mbox{if } m = 0 \\\\ A(m-1, 1) & \\mbox{if } m > 0 \\mbox{ and } n = 0 \\\\ A(m-1, A(m, n-1)) & \\mbox{if } m > 0 \\mbox{ and } n > 0. \\end{cases}$

معلماتها لا تكون أبدًا سالبة وهي دائمًا تنتهي.

# --instructions--

اكتب دالة تُرجع قيمة $A(m, n)$. يُفضل استخدام دقة عشوائية (نظرًا لأن الدالة تنمو بسرعة كبيرة)، لكنها ليست مطلوبة.

# --hints--

`ack` يجب أن تكون دالة.

```js
assert(typeof ack === 'function');
```

`ack(0, 0)` يجب أن تُرجع 1.

```js
assert(ack(0, 0) === 1);
```

`ack(1, 1)` يجب أن تُرجع 3.

```js
assert(ack(1, 1) === 3);
```

`ack(2, 5)` يجب أن تُرجع 13.

```js
assert(ack(2, 5) === 13);
```

`ack(3, 3)` يجب أن تُرجع 61.

```js
assert(ack(3, 3) === 61);
```

# --seed--

## --seed-contents--

```js
function ack(m, n) {

}
```

# --solutions--

```js
function ack(m, n) {
  return m === 0 ? n + 1 : ack(m - 1, n === 0 ? 1 : ack(m, n - 1));
}
```
