---
id: 5949b579404977fbaefcd736
title: 9 مليارات أسماء الله العدد الصحيح
challengeType: 1
forumTopicId: 302219
dashedName: 9-billion-names-of-god-the-integer
---

# --description--

هذه المهمة هي تعديل لقصة قصيرة لآرثر سي. كلارك.

(يجب أن يكون الحلّالون على دراية بعواقب إتمام هذه المهمة.)

بالتفصيل، لتحديد المقصود بـ "اسم":

<ul>
<li>العدد الصحيح 1 له اسم واحد "1".</li>
  <li>العدد الصحيح 2 له اسمان "1+1" و"2".</li>
  <li>العدد الصحيح 3 له ثلاثة أسماء "1+1+1"، "2+1"، و"3".</li>
  <li>العدد الصحيح 4 له خمسة أسماء "1+1+1+1"، "2+1+1"، "2+2"، "3+1"، "4".</li>
  <li>العدد الصحيح 5 له سبعة أسماء "1+1+1+1+1"، "2+1+1+1"، "2+2+1"، "3+1+1"، "3+2"، "4+1"، "5".</li>
</ul>

يمكن تصور ذلك بالشكل التالي:

<pre>          1
        1   1
      1   1   1
    1   2   1   1
  1   2   2   1   1
1   3   3   2   1   1
</pre>

حيث الصف $n$ يقابل العدد الصحيح $n$، وكل عمود $C$ في الصف $m$ من اليسار إلى اليمين يقابل عدد الأسماء التي تبدأ بـ $C$.

يمكن ملاحظة أن مجموع الصف $n$ $P(n)$ هو دالة تقسيم العدد الصحيح.

# --instructions--

نفذ دالة تُرجع مجموع الصف $n$.

# --hints--

يجب أن تكون `numberOfNames` دالة.

```js
assert(typeof numberOfNames === 'function');
```

يجب أن تساوي `numberOfNames(5)` القيمة 7.

```js
assert.equal(numberOfNames(5), 7);
```

يجب أن تساوي `numberOfNames(12)` القيمة 77.

```js
assert.equal(numberOfNames(12), 77);
```

يجب أن تساوي `numberOfNames(18)` القيمة 385.

```js
assert.equal(numberOfNames(18), 385);
```

يجب أن تساوي `numberOfNames(23)` القيمة 1255.

```js
assert.equal(numberOfNames(23), 1255);
```

يجب أن تساوي `numberOfNames(42)` القيمة 53174.

```js
assert.equal(numberOfNames(42), 53174);
```

يجب أن تساوي `numberOfNames(123)` القيمة 2552338241.

```js
assert.equal(numberOfNames(123), 2552338241);
```

# --seed--

## --seed-contents--

```js
function numberOfNames(num) {

  return true;
}
```

# --solutions--

```js
function numberOfNames(num) {
  const cache = [
    [1]
  ];
  for (let l = cache.length; l < num + 1; l++) {
    let Aa;
    let Mi;
    const r = [0];
    for (let x = 1; x < l + 1; x++) {
      r.push(r[r.length - 1] + (Aa = cache[l - x < 0 ? cache.length - (l - x) : l - x])[(Mi = Math.min(x, l - x)) < 0 ? Aa.length - Mi : Mi]);
    }
    cache.push(r);
  }
  return cache[num][cache[num].length - 1];
}
```
