---
id: 594dc6c729e5700999302b45
title: أقواس متوازنة
challengeType: 1
forumTopicId: 302230
dashedName: balanced-brackets
---

# --description--

حدد ما إذا كانت سلسلة الأقواس المولدة متوازنة؛ أي ما إذا كانت تتكون بالكامل من أزواج من أقواس الفتح/الإغلاق (بهذا الترتيب)، دون وجود أي تسكين خاطئ.

**أمثلة:**
| المُدخَل                     | الناتج |
| ------------------------- | ------ |
| <code>\[]</code>          | true   |
| <code>]\[</code>          | false  |
| <code>[][]</code>         | true   |
| <code>]\[]</code>         | false  |
| <code>\[]]\[\[]</code>    | false  |
| <code>\[\[\[\[]]]]</code> | true   |

# --before-each--

```js
const testCases = [
  '[]',
  ']][[[][][][]][',
  '[][[[[][][[[]]]]]]',
  '][',
  '[[[]]]][[]',
  '][[]',
  '][[][]][[[]]',
  '[[][]]][',
  '[[[]]][[]]]][][[',
  '[]][[]]][[[[][]]',
  '][]][[][',
  '[[]][[][]]',
  '[[]]',
  ']][]][[]][[[',
  '][]][][[',
  '][][',
  '[]]]',
  ''
];
```

# --hints--

يجب أن تكون `isBalanced` دالة.

```js
assert(typeof isBalanced === 'function');
```

يجب أن تُرجع `isBalanced("[]")` القيمة true.

```js
assert(isBalanced(testCases[0]));
```

يجب أن تُرجع `isBalanced("]][[[][][][]][")` القيمة false.

```js
assert(!isBalanced(testCases[1]));
```

يجب أن تُرجع `isBalanced("[][[[[][][[[]]]]]]")` القيمة true.

```js
assert(isBalanced(testCases[2]));
```

يجب أن تُرجع `isBalanced("][")` القيمة false.

```js
assert(!isBalanced(testCases[3]));
```

يجب أن تُرجع `isBalanced("[[[]]]][[]")` القيمة false.

```js
assert(!isBalanced(testCases[4]));
```

يجب أن تُرجع `isBalanced("][[]")` القيمة false.

```js
assert(!isBalanced(testCases[5]));
```

يجب أن تُرجع `isBalanced("][[][]][[[]]")` القيمة false.

```js
assert(!isBalanced(testCases[6]));
```

يجب أن تُرجع `isBalanced("[[][]]][")` القيمة false.

```js
assert(!isBalanced(testCases[7]));
```

يجب أن تُرجع `isBalanced("[[[]]][[]]]][][[")` القيمة false.

```js
assert(!isBalanced(testCases[8]));
```

يجب أن تُرجع `isBalanced("[]][[]]][[[[][]]")` القيمة false.

```js
assert(!isBalanced(testCases[9]));
```

يجب أن تُرجع `isBalanced("][]][[][")` القيمة false.

```js
assert(!isBalanced(testCases[10]));
```

يجب أن تُرجع `isBalanced("[[]][[][]]")` القيمة true.

```js
assert(isBalanced(testCases[11]));
```

يجب أن تُرجع `isBalanced("[[]]")` القيمة true.

```js
assert(isBalanced(testCases[12]));
```

يجب أن تُرجع `isBalanced("]][]][[]][[[")` القيمة false.

```js
assert(!isBalanced(testCases[13]));
```

يجب أن تُرجع `isBalanced("][]][][[")` القيمة false.

```js
assert(!isBalanced(testCases[14]));
```

يجب أن تُرجع `isBalanced("][][")` القيمة false.

```js
assert(!isBalanced(testCases[15]));
```

يجب أن تُرجع `isBalanced("[]]]")` القيمة false.

```js
assert(!isBalanced(testCases[16]));
```

يجب أن تُرجع `isBalanced("")` القيمة true.

```js
assert(isBalanced(testCases[17]));
```

# --seed--

## --seed-contents--

```js
function isBalanced(str) {

  return true;
}
```

# --solutions--

```js
function isBalanced(str) {
  if (str === '') return true;
  let a = str;
  let b;
  do {
    b = a;
    a = a.replace(/\[\]/g, '');
  } while (a !== b);
  return !a;
}
```
