---
id: 594da033de4190850b893874
title: المتوسطات/الجذر التربيعي المتوسط
challengeType: 1
forumTopicId: 302228
dashedName: averagesroot-mean-square
---

# --description--

احسب الجذر التربيعي المتوسط (RMS) للأعداد من 1 إلى 10 شاملة.

<abbr title="الجذر التربيعي المتوسط">RMS</abbr> يُحسب بأخذ الجذر التربيعي لمتوسط مربعات الأعداد، كما هو معطى في المعادلة:

$$x\_{\\mathrm{rms}} = \\sqrt {{{x_1}^2 + {x_2}^2 + \\cdots + {x_n}^2} \\over n}. $$

# --hints--

يجب أن تكون `rms` دالة.

```js
assert(typeof rms === 'function');
```

يجب أن تساوي `rms([1, 2, 3, 4, 5, 6, 7, 8, 9, 10])` القيمة `6.2048368229954285`.

```js
const arr1 = [1, 2, 3, 4, 5, 6, 7, 8, 9, 10];
const answer1 = 6.2048368229954285;
assert.equal(rms(arr1), answer1);
```

# --seed--

## --seed-contents--

```js
function rms(arr) {

}
```

# --solutions--

```js
function rms(arr) {
  const sumOfSquares = arr.reduce((s, x) => s + x * x, 0);
  return Math.sqrt(sumOfSquares / arr.length);
}
```
