---
id: 594ecc0d9a8cf816e3340187
title: أعداد التاكسي
challengeType: 1
forumTopicId: 302337
dashedName: taxicab-numbers
---

# --description--

عدد التاكسي (التعريف المستخدم هنا) هو عدد صحيح موجب يمكن التعبير عنه كمجموع مكعبين موجبين بطريقتين أو أكثر.

أول عدد تاكسي هو `1729`، وهو:

1<sup>3</sup> + 12<sup>3</sup> و

9<sup>3</sup> + 10<sup>3</sup>.

تُعرف أعداد التاكسي أيضًا باسم:

<ul>
<li>أعداد التاكسي</li>
  <li>أعداد تاكسي-تاب</li>
  <li>أعداد تاكسي تاب</li>
  <li>أعداد هاردي-رامانوجان</li>
</ul>

# --instructions--

اكتب دالة تُرجع أقل `n` أعداد تاكسي. لكل عدد تاكسي، اعرض العدد بالإضافة إلى مكعباته المكونة له.

# --before-each--

```js
const res4 = [1729, 4104, 13832, 20683];
const res25 = [
  1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656,
  110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763,
  373464, 402597
];

const res39From20To29 = [314496, 320264, 327763, 373464, 402597, 439101, 443889, 513000, 513856];
```

# --hints--

يجب أن تكون `taxicabNumbers` دالة.

```js
assert(typeof taxicabNumbers === 'function');
```

يجب أن تُرجع `taxicabNumbers` مصفوفة.

```js
assert(typeof taxicabNumbers(2) === 'object');
```

يجب أن تُرجع `taxicabNumbers` مصفوفة من الأعداد.

```js
assert(typeof taxicabNumbers(100)[0] === 'number');
```

يجب أن تُرجع `taxicabNumbers(4)` [1729, 4104, 13832, 20683].

```js
assert.deepEqual(taxicabNumbers(4), res4);
```

يجب أن تُرجع `taxicabNumbers(25)` [1729, 4104, 13832, 20683, 32832, 39312, 40033, 46683, 64232, 65728, 110656, 110808, 134379, 149389, 165464, 171288, 195841, 216027, 216125, 262656, 314496, 320264, 327763, 373464, 402597].

```js
assert.deepEqual(taxicabNumbers(25), res25);
```

يجب أن تكون الأعداد الناتجة من 20 إلى 29 في `taxicabNumbers(39)` هي [314496,320264,327763,373464,402597,439101,443889,513000,513856].

```js
assert.deepEqual(taxicabNumbers(39).slice(20, 29), res39From20To29);
```

# --seed--

## --seed-contents--

```js
function taxicabNumbers(n) {

  return true;
}
```

# --solutions--

```js
function taxicabNumbers(nNumbers) {
  const cubeN = [];
  const s3s = {};

  const e = 100;
  for (let n = 1; n < e; n += 1) {
    cubeN[n] = n * n * n;
  }

  for (let a = 1; a < e - 1; a += 1) {
    const a3 = cubeN[a];
    for (let b = a; b < e; b += 1) {
      const b3 = cubeN[b];
      const s3 = a3 + b3;

      let abs = s3s[s3];
      if (!abs) {
        s3s[s3] = abs = [];
      }
      abs.push([a, b]);
    }
  }

  let i = 0;
  const res = [];
  Object.keys(s3s).forEach(s3 => {
    const abs = s3s[s3];
    if (abs.length >= 2) { // No two cube pairs found
      i += 1;
      if (i <= nNumbers) {
        res.push(s3);
      }
    }
  });
  return res.map(item => parseInt(item, 10));
}
```
