---
id: 594fa2746886f41f7d8bf225
title: الفرز الطوبولوجي
challengeType: 1
forumTopicId: 302340
dashedName: topological-sort
---

# --description--

بما أن هناك خريطة بين العناصر والعناصر التي تعتمد عليها، يرتب الفرز الطوبولوجي العناصر بحيث لا يسبق أي عنصر عنصرًا يعتمد عليه. هناك خوارزميتان شائعتان للفرز الطوبولوجي: فرز كاهن (1962) والبحث بالعمق أولاً.

# --instructions--

اكتب دالة تُرجع قائمة بترتيب تجميع صحيح للمكتبات بناءً على تبعياتها.

- افترض أن أسماء المكتبات كلمة واحدة.
- العناصر المذكورة كمعتمدة فقط لا تملك معتمدين خاصين بها، لكن يجب تحديد ترتيب تجميعها.
- تجاهل أي تبعيات ذاتية.
- تجاهل أي تبعيات لا يمكن ترتيبها.

استخدم البيانات التالية كمثال:

<pre>
المكتبة          تبعيات المكتبة
=======          ====================
des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
dw01             ieee dw01 dware gtech
dw02             ieee dw02 dware
dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
dw04             dw04 ieee dw01 dware gtech
dw05             dw05 ieee dware
dw06             dw06 ieee dware
dw07             ieee dware
dware            ieee dware
gtech            ieee gtech
ramlib           std ieee
std_cell_lib     ieee std_cell_lib
synopsys
</pre>

تجميع مكتبة في لغة VHDL يخضع لقاعدة أن المكتبة يجب أن تُجمع بعد أي مكتبة تعتمد عليها. ستكون البيانات أعلاه غير قابلة للترتيب إذا، على سبيل المثال، أُضيف `dw04` إلى قائمة تبعيات `dw01`.

سيكون مُدخل الدالة سلسلة نصية متعددة الأسطر، كل سطر يحتوي على اسم المكتبة متبوعًا بتبعياتها (إن وجدت).

على سبيل المثال:

```js
const libsSimple =
  `aaa bbb
  bbb`;
```

# --before-each--

```js
const libsSimple =
  `aaa bbb
  bbb`;

const libsVHDL =
  `des_system_lib   std synopsys std_cell_lib des_system_lib dw02 dw01 ramlib ieee
  dw01             ieee dw01 dware gtech
  dw02             ieee dw02 dware
  dw03             std synopsys dware dw03 dw02 dw01 ieee gtech
  dw04             dw04 ieee dw01 dware gtech
  dw05             dw05 ieee dware
  dw06             dw06 ieee dware
  dw07             ieee dware
  dware            ieee dware
  gtech            ieee gtech
  ramlib           std ieee
  std_cell_lib     ieee std_cell_lib
  synopsys`;

const libsCustom =
  `a b c d
  b c d
  d c
  c base
  base`;

const libsUnorderable =
  `TestLib Base MainLib
  MainLib TestLib
  Base`;
```

# --hints--

يجب أن تكون `topologicalSort` دالة.

```js
assert(typeof topologicalSort === 'function');
```

يجب أن تُرجع `topologicalSort(libsSimple)` مصفوفة.

```js
assert(Array.isArray(topologicalSort(libsSimple)));
```

يجب أن تُرجع `topologicalSort(libsSimple)` القيمة `['bbb', 'aaa']`.

```js
assert.deepEqual(topologicalSort(libsSimple), ['bbb', 'aaa']);
```

يجب أن تُرجع `topologicalSort(libsVHDL)` القيمة `['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']`.

```js
assert.deepEqual(topologicalSort(libsVHDL), ['ieee', 'std_cell_lib', 'gtech', 'dware', 'dw07', 'dw06', 'dw05', 'dw02', 'dw01', 'dw04', 'std', 'ramlib', 'synopsys', 'dw03', 'des_system_lib']);
```

يجب أن تُرجع `topologicalSort(libsCustom)` القيمة `['base', 'c', 'd', 'b', 'a']`.

```js
assert.deepEqual(topologicalSort(libsCustom), ['base', 'c', 'd', 'b', 'a']);
```

يجب أن تتجاهل `topologicalSort` التبعيات التي لا يمكن ترتيبها.

```js
assert.deepEqual(topologicalSort(libsUnorderable), ['Base']);
```

# --seed--

## --seed-contents--

```js
function topologicalSort(libs) {

  return true;
}
```

# --solutions--

```js
function topologicalSort(libs) {
  // A map of the input data, with the keys as the packages, and the values as
  // and array of packages on which it depends.
  const D = libs
    .split('\n')
    .map(e => e.split(' ').filter(ep => ep !== ''))
    .reduce((p, c) =>
      p.set(c[0], c.filter((e, i) => (i > 0 && e !== c[0] ? e : null))), new Map());
  [].concat(...D.values()).forEach(e => {
    D.set(e, D.get(e) || []);
  });

  // The above map rotated so that it represents a DAG of the form
  // Map {
  //    A => [ A, B, C],
  //    B => [C],
  //    C => []
  // }
  // where each key represents a node, and the array contains the edges.
  const G = [...D.keys()].reduce((p, c) =>
    p.set(
      c,
      [...D.keys()].filter(e => D.get(e).includes(c))),
    new Map()
  );

  // An array of leaf nodes; nodes with 0 in degrees.
  const Q = [...D.keys()].filter(e => D.get(e).length === 0);

  // The result array.
  const S = [];
  while (Q.length) {
    const u = Q.pop();
    S.push(u);
    G.get(u).forEach(v => {
      D.set(v, D.get(v).filter(e => e !== u));
      if (D.get(v).length === 0) {
        Q.push(v);
      }
    });
  }

  return S;
}
```
