---
id: 594faaab4e2a8626833e9c3d
title: تقسيم سلسلة نصية مع الهروب
challengeType: 1
forumTopicId: 302338
dashedName: tokenize-a-string-with-escaping
---

# --description--

اكتب دالة أو برنامجًا يمكنه تقسيم سلسلة نصية عند كل ظهور لفاصل لم يتم الهروب منه.

يجب أن تقبل ثلاثة معلمات مُدخَل:

<ul>
<li>السلسلة النصية</li>
  <li>حرف الفاصل</li>
  <li>حرف الهروب</li>
</ul>

يجب أن تُخرج قائمة من السلاسل النصية.

قواعد التقسيم:

<ul>
<li>الحقول التي فُصلت بواسطة الفواصل تصبح عناصر في قائمة الإخراج.</li>
  <li>يجب الحفاظ على الحقول الفارغة، حتى في البداية والنهاية.</li>
</ul>

قواعد الهروب:

<ul>
<li>"الهروب" يعني أن يسبقه ظهور لحرف الهروب لم يتم الهروب منه مسبقًا.</li>
  <li>عندما يسبق حرف الهروب حرفًا لا يحمل معنى خاصًا، يُعتبر هروبًا (لكن لا يفعل شيئًا خاصًا).</li>
  <li>كل ظهور لحرف الهروب الذي استُخدم للهروب من شيء ما، لا يجب أن يصبح جزءًا من الإخراج.</li>
</ul>

أثبت أن دالتك تحقق حالة الاختبار التالية:

بالنظر إلى السلسلة النصية

<pre>one^|uno||three^^^^|four^^^|^cuatro|</pre>

وباستخدام `|` كفاصل و`^` كحرف هروب، يجب أن تُخرج دالتك المصفوفة التالية:

<pre>  ['one|uno', '', 'three^^', 'four^|cuatro', '']
</pre>

# --before-each--

```js
const testStr1 = 'one^|uno||three^^^^|four^^^|^cuatro|';
const res1 = ['one|uno', '', 'three^^', 'four^|cuatro', ''];

const testStr2 = 'a@&bcd&ef&&@@hi';
const res2 = ['a&bcd', 'ef', '', '@hi'];

const testStr3 = 'hello^|world|how^are^you^|';
const res3 = ['hello|world', 'howareyou|'];
```

# --hints--

يجب أن تكون `tokenize` دالة.

```js
assert(typeof tokenize === 'function');
```

يجب أن تُرجع `tokenize` مصفوفة.

```js
assert(typeof tokenize('a', 'b', 'c') === 'object');
```

`tokenize('one^|uno||three^^^^|four^^^|^cuatro|', '|', '^')` يجب أن تُرجع `['one|uno', '', 'three^^', 'four^|cuatro', '']`

```js
assert.deepEqual(tokenize(testStr1, '|', '^'), res1);
```

`tokenize('a@&bcd&ef&&@@hi', '&', '@')` يجب أن تُرجع `['a&bcd', 'ef', '', '@hi']`

```js
assert.deepEqual(tokenize(testStr2, '&', '@'), res2);
```

`tokenize('hello^|world|how^are^you^|', '|', '^')` يجب أن تُرجع `['hello|world', 'howareyou|']`

```js
assert.deepEqual(tokenize(testStr3, '|', '^'), res3);
```

# --seed--

## --seed-contents--

```js
function tokenize(str, sep, esc) {
  return true;
}
```

# --solutions--

```js
// tokenize :: String -> Character -> Character -> [String]
function tokenize(str, charDelim, charEsc) {
  const dctParse = str.split('')
    .reduce((a, x) => {
      const blnEsc = a.esc;
      const blnBreak = !blnEsc && x === charDelim;
      const blnEscChar = !blnEsc && x === charEsc;

      return {
        esc: blnEscChar,
        token: blnBreak ? '' : (
          a.token + (blnEscChar ? '' : x)
        ),
        list: a.list.concat(blnBreak ? a.token : [])
      };
    }, {
      esc: false,
      token: '',
      list: []
    });

  return dctParse.list.concat(
    dctParse.token
  );
}
```
