---
id: 5951815dd895584b06884620
title: دوائر بنصف قطر معين تمر بنقطتين
challengeType: 1
forumTopicId: 302231
dashedName: circles-of-given-radius-through-two-points
---

# --description--

عند وجود نقطتين على مستوى ونصف قطر، عادةً يمكن رسم دائرتين بنصف القطر المعطى تمران بتلك النقطتين.

**استثناءات:**

<ul>
<li>يجب اعتبار نصف القطر صفر على أنه لا يصف دوائر أبدًا (باستثناء حالة تزامن النقطتين).</li>
  <li>إذا كانت النقطتان متطابقتين، يمكن رسم عدد لا نهائي من الدوائر التي تمر بالنقطة على محيطها، إلا إذا كان نصف القطر صفرًا أيضًا، حينها تنهار الدوائر إلى نقطة واحدة.</li>
  <li>إذا شكلت النقطتان قطرًا، فارجع بدائرة واحدة فقط.</li>
  <li>إذا كانت النقطتان بعيدتين جدًا، فلا يمكن رسم أي دوائر.</li>
</ul>

# --instructions--

نفذ دالة تأخذ نقطتين ونصف قطر وتُرجع الدائرتين اللتين تمران بتلك النقطتين. لكل دائرة ناتجة، قدّم إحداثيات مركز الدائرة مقربة إلى أربعة أرقام عشرية. أرجع كل إحداثية كمصفوفة، والإحداثيات كمصفوفة من المصفوفات.

**لحالات الحافة، أرجع التالي:**

<ul>
<li>إذا كانت النقطتان على القطر، أرجع نقطة واحدة. إذا كان نصف القطر صفرًا أيضًا، أرجع <code>"Radius Zero"</code>.</li>
  <li>إذا كانت النقطتان متطابقتين، أرجع <code>"Coincident point. Infinite solutions"</code>.</li>
  <li>إذا كانت النقطتان أبعد من القطر، أرجع <code>"No intersection. Points further apart than circle diameter"</code>.</li>
</ul>

**مدخلات نموذجية:**

<pre>      p1                p2           r
0.1234, 0.9876    0.8765, 0.2345    2.0
0.0000, 2.0000    0.0000, 0.0000    1.0
0.1234, 0.9876    0.1234, 0.9876    2.0
0.1234, 0.9876    0.8765, 0.2345    0.5
0.1234, 0.9876    0.1234, 0.9876    0.0
</pre>

# --before-each--

```js
const testCases = [
  [[0.1234, 0.9876], [0.8765, 0.2345], 2.0],
  [[0.0000, 2.0000], [0.0000, 0.0000], 1.0],
  [[0.1234, 0.9876], [0.1234, 0.9876], 2.0],
  [[0.1234, 0.9876], [0.8765, 0.2345], 0.5],
  [[0.1234, 0.9876], [0.1234, 0.9876], 0.0]
];
const answers = [
  [[1.8631, 1.9742], [-0.8632, -0.7521]],
  [0, 1],
  'Coincident point. Infinite solutions',
  'No intersection. Points further apart than circle diameter',
  'Radius Zero'
];
```

# --hints--

يجب أن تكون `getCircles` دالة.

```js
assert(typeof getCircles === 'function');
```

يجب أن تُرجع `getCircles([0.1234, 0.9876], [0.8765, 0.2345], 2.0)` القيمة `[[1.8631, 1.9742], [-0.8632, -0.7521]]`.

```js
assert.deepEqual(getCircles(...testCases[0]), answers[0]);
```

يجب أن تُرجع `getCircles([0.0000, 2.0000], [0.0000, 0.0000], 1.0)` القيمة `[0, 1]`

```js
assert.deepEqual(getCircles(...testCases[1]), answers[1]);
```

يجب أن تُرجع `getCircles([0.1234, 0.9876], [0.1234, 0.9876], 2.0)` القيمة `Coincident point. Infinite solutions`

```js
assert.deepEqual(getCircles(...testCases[2]), answers[2]);
```

يجب أن تُرجع `getCircles([0.1234, 0.9876], [0.8765, 0.2345], 0.5)` القيمة `No intersection. Points further apart than circle diameter`

```js
assert.deepEqual(getCircles(...testCases[3]), answers[3]);
```

يجب أن تُرجع `getCircles([0.1234, 0.9876], [0.1234, 0.9876], 0.0)` القيمة `Radius Zero`

```js
assert.deepEqual(getCircles(...testCases[4]), answers[4]);
```

# --seed--

## --seed-contents--

```js
function getCircles(...args) {

  return true;
}
```

# --solutions--

```js
const hDist = (p1, p2) => Math.hypot(...p1.map((e, i) => e - p2[i])) / 2;
const pAng = (p1, p2) => Math.atan(p1.map((e, i) => e - p2[i]).reduce((p, c) => c / p, 1));
const solveF = (p, r) => t => [parseFloat((r * Math.cos(t) + p[0]).toFixed(4)), parseFloat((r * Math.sin(t) + p[1]).toFixed(4))];
const diamPoints = (p1, p2) => p1.map((e, i) => parseFloat((e + (p2[i] - e) / 2).toFixed(4)));

function getCircles(...args) {
  const [p1, p2, s] = args;
  const solve = solveF(p1, s);
  const halfDist = hDist(p1, p2);

  let msg = [];
  switch (Math.sign(s - halfDist)) {
    case 0:
      msg = s ? diamPoints(p1, p2) :
        'Radius Zero';
      break;
    case 1:
      if (!halfDist) {
        msg = 'Coincident point. Infinite solutions';
      }
      else {
        const theta = pAng(p1, p2);
        const theta2 = Math.acos(halfDist / s);
        [1, -1].map(e => solve(theta + e * theta2)).forEach(
          e => msg.push(e));
      }
      break;
    case -1:
      msg = 'No intersection. Points further apart than circle diameter';
      break;
    default:
      msg = 'Reached the default';
  }
  return msg;
}
```
