---
id: 5956795bc9e2c415eb244de1
title: الربط باستخدام التجزئة
challengeType: 1
forumTopicId: 302284
dashedName: hash-join
---

# --description--

الربط الداخلي هو عملية تجمع بين جدولين من البيانات في جدول واحد، بناءً على تطابق قيم الأعمدة. أبسط طريقة لتنفيذ هذه العملية هي خوارزمية الربط بالتكرار المتداخل، لكن بديلًا أكثر قابلية للتوسع هو خوارزمية الربط باستخدام التجزئة.

تتكون خوارزمية "الربط باستخدام التجزئة" من خطوتين:

<ol>
<li><strong>مرحلة التجزئة:</strong> أنشئ خريطة متعددة من أحد الجدولين، تربط كل قيمة في عمود الربط بكل الصفوف التي تحتويها.</li>
  <ul>
<li>يجب أن تدعم الخريطة المتعددة البحث باستخدام التجزئة الذي يتوسع بشكل أفضل من البحث الخطي البسيط، لأن هذه هي الفكرة الأساسية لهذه الخوارزمية.</li>
    <li>من الأفضل إنشاء الخريطة المتعددة للجدول الأصغر، لتقليل وقت إنشائها وحجم الذاكرة المستخدمة.</li>
  </ul>
<li><strong>مرحلة الربط:</strong> امسح الجدول الآخر، وابحث عن الصفوف المطابقة من خلال النظر في الخريطة المتعددة التي أنشئت سابقًا.</li>
</ol>

يمكن التعبير عن الخوارزمية بالشيفرة الزائفة كما يلي:

<pre><strong>let</strong> <i>A</i> = الجدول الأول المُدخل (أو الأفضل، الأكبر)
<strong>let</strong> <i>B</i> = الجدول الثاني المُدخل (أو الأفضل، الأصغر)
<strong>let</strong> <i>j<sub>A</sub></i> = معرف عمود الربط في الجدول <i>A</i>
<strong>let</strong> <i>j<sub>B</sub></i> = معرف عمود الربط في الجدول <i>B</i>
<strong>let</strong> <i>M<sub>B</sub></i> = خريطة متعددة تربط القيم المفردة بعدة صفوف من الجدول <i>B</i> (تبدأ فارغة)
<strong>let</strong> <i>C</i> = جدول الإخراج (يبدأ فارغًا)
<strong>for each</strong> صف <i>b</i> في الجدول <i>B</i>:
  <strong>place</strong> <i>b</i> في الخريطة المتعددة <i>M<sub>B</sub></i> تحت المفتاح <i>b(j<sub>B</sub>)</i>
<strong>for each</strong> صف <i>a</i> في الجدول <i>A</i>:
  <strong>for each</strong> صف <i>b</i> في الخريطة المتعددة <i>M<sub>B</sub></i> تحت المفتاح <i>a(j<sub>A</sub>)</i>:
    <strong>let</strong> <i>c</i> = دمج الصف <i>a</i> مع الصف <i>b</i>
    <strong>place</strong> الصف <i>c</i> في الجدول <i>C</i>
</pre>

# --instructions--

نفذ خوارزمية "الربط باستخدام التجزئة" كدالة وأثبت أنها تجتاز حالة الاختبار المذكورة أدناه. يجب أن تقبل الدالة مصفوفتين من الكائنات وتُرجع مصفوفة من الكائنات المدمجة.

**المدخلات**

<table>
  <tr>
    <td style="padding: 4px; margin: 5px;">
      <table style="border:none; border-collapse:collapse;">
        <tr>
          <td style="border:none"><i>A =</i></td>
          <td style="border:none">
            <table>
              <tr>
<th style="padding: 4px; margin: 5px;">العمر</th>
                <th style="padding: 4px; margin: 5px;">الاسم</th>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">27</td>
<td style="padding: 4px; margin: 5px;">Jonah</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">18</td>
<td style="padding: 4px; margin: 5px;">Alan</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">28</td>
<td style="padding: 4px; margin: 5px;">Glory</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">18</td>
<td style="padding: 4px; margin: 5px;">Popeye</td>
              </tr>
              <tr>
                <td style="padding: 4px; margin: 5px;">28</td>
<td style="padding: 4px; margin: 5px;">Alan</td>
              </tr>
            </table>
          </td>
          <td style="border:none; padding-left:1.5em;" rowspan="2"></td>
          <td style="border:none"><i>B =</i></td>
          <td style="border:none">
            <table>
              <tr>
<th style="padding: 4px; margin: 5px;">الشخصية</th>
                <th style="padding: 4px; margin: 5px;">الخصم</th>
              </tr>
              <tr>
<td style="padding: 4px; margin: 5px;">Jonah</td>
                <td style="padding: 4px; margin: 5px;">الحيتان</td>
              </tr>
              <tr>
<td style="padding: 4px; margin: 5px;">Jonah</td>
                <td style="padding: 4px; margin: 5px;">العناكب</td>
              </tr>
              <tr>
<td style="padding: 4px; margin: 5px;">Alan</td>
                <td style="padding: 4px; margin: 5px;">الأشباح</td>
              </tr>
              <tr>
<td style="padding: 4px; margin: 5px;">Alan</td>
                <td style="padding: 4px; margin: 5px;">الزومبي</td>
              </tr>
              <tr>
<td style="padding: 4px; margin: 5px;">Glory</td>
                <td style="padding: 4px; margin: 5px;">بافي</td>
              </tr>
            </table>
          </td>
        </tr>
        <tr>
          <td style="border:none">
            <i>j<sub>A</sub> =</i>
          </td>
          <td style="border:none">
<i><code>Name</code> (أي العمود 1)</i>
          </td>
          <td style="border:none">
            <i>j<sub>B</sub> =</i>
          </td>
          <td style="border:none">
<i><code>Character</code> (أي العمود 0)</i>
          </td>
        </tr>
      </table>
    </td>
  </tr>
</table>

**المخرجات**

| A_age | A_name | B_character | B_nemesis |
| ----- | ------ | ----------- | --------- |
| 27    | Jonah  | Jonah       | الحيتان    |
| 27    | Jonah  | Jonah       | العناكب   |
| 18    | Alan   | Alan        | الأشباح    |
| 18    | Alan   | Alan        | الزومبي   |
| 28    | Glory  | Glory       | بافي     |
| 28    | Alan   | Alan        | الأشباح    |
| 28    | Alan   | Alan        | الزومبي   |

ترتيب الصفوف في جدول الإخراج غير مهم.

# --hints--

يجب أن تكون `hashJoin` دالة.

```js
assert(typeof hashJoin === 'function');
```

`hashJoin([{ age: 27, name: "Jonah" }, { age: 18, name: "Alan" }, { age: 28, name: "Glory" }, { age: 18, name: "Popeye" }, { age: 28, name: "Alan" }], [{ character: "Jonah", nemesis: "Whales" }, { character: "Jonah", nemesis: "Spiders" }, { character: "Alan", nemesis: "Ghosts" }, { character:"Alan", nemesis: "Zombies" }, { character: "Glory", nemesis: "Buffy" }, { character: "Bob", nemesis: "foo" }])` يجب أن تُرجع `[{"A_age": 27,"A_name": "Jonah", "B_character": "Jonah", "B_nemesis": "Whales"}, {"A_age": 27,"A_name": "Jonah", "B_character": "Jonah", "B_nemesis": "Spiders"}, {"A_age": 18,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Ghosts"}, {"A_age": 18,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Zombies"}, {"A_age": 28,"A_name": "Glory", "B_character": "Glory", "B_nemesis": "Buffy"}, {"A_age": 28,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Ghosts"}, {"A_age": 28,"A_name": "Alan", "B_character": "Alan", "B_nemesis": "Zombies"}]`

```js
const hash1 = [
    { age: 27, name: 'Jonah' },
    { age: 18, name: 'Alan' },
    { age: 28, name: 'Glory' },
    { age: 18, name: 'Popeye' },
    { age: 28, name: 'Alan' }
];

const hash2 = [
    { character: 'Jonah', nemesis: 'Whales' },
    { character: 'Jonah', nemesis: 'Spiders' },
    { character: 'Alan', nemesis: 'Ghosts' },
    { character: 'Alan', nemesis: 'Zombies' },
    { character: 'Glory', nemesis: 'Buffy' },
    { character: 'Bob', nemesis: 'foo' }
];

const res = [
    { A_age: 27, A_name: 'Jonah', B_character: 'Jonah', B_nemesis: 'Whales' },
    { A_age: 27, A_name: 'Jonah', B_character: 'Jonah', B_nemesis: 'Spiders' },
    { A_age: 18, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Ghosts' },
    { A_age: 18, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Zombies' },
    { A_age: 28, A_name: 'Glory', B_character: 'Glory', B_nemesis: 'Buffy' },
    { A_age: 28, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Ghosts' },
    { A_age: 28, A_name: 'Alan', B_character: 'Alan', B_nemesis: 'Zombies' }
];

const bench1 = [{ name: 'u2v7v', num: 1 }, { name: 'n53c8', num: 10 }, { name: 'oysce', num: 9 }, { name: '0mto2s', num: 1 }, { name: 'vkh5id', num: 4 }, { name: '5od0cf', num: 8 }, { name: 'uuulue', num: 10 }, { name: '3rgsbi', num: 9 }, { name: 'kccv35r', num: 4 }, { name: '80un74', num: 9 }, { name: 'h4pp3', num: 6 }, { name: '51bit', num: 7 }, { name: 'j9ndf', num: 8 }, { name: 'vf3u1', num: 10 }, { name: 'g0bw0om', num: 10 }, { name: 'j031x', num: 7 }, { name: 'ij3asc', num: 9 }, { name: 'byv83y', num: 8 }, { name: 'bjzp4k', num: 4 }, { name: 'f3kbnm', num: 10 }];
const bench2 = [{ friend: 'o8b', num: 8 }, { friend: 'ye', num: 2 }, { friend: '32i', num: 5 }, { friend: 'uz', num: 3 }, { friend: 'a5k', num: 4 }, { friend: 'uad', num: 7 }, { friend: '3w5', num: 10 }, { friend: 'vw', num: 10 }, { friend: 'ah', num: 4 }, { friend: 'qv', num: 7 }, { friend: 'ozv', num: 2 }, { friend: '9ri', num: 10 }, { friend: '7nu', num: 4 }, { friend: 'w3', num: 9 }, { friend: 'tgp', num: 8 }, { friend: 'ibs', num: 1 }, { friend: 'ss7', num: 6 }, { friend: 'g44', num: 9 }, { friend: 'tab', num: 9 }, { friend: 'zem', num: 10 }];
assert.deepEqual(hashJoin(hash1, hash2), res);
```

# --seed--

## --seed-contents--

```js
function hashJoin(hash1, hash2) {

  return [];
}
```

# --solutions--

```js
function hashJoin(hash1, hash2) {
  const hJoin = (tblA, tblB, strJoin) => {
    const [jA, jB] = strJoin.split('=');
    const M = tblB.reduce((a, x) => {
      const id = x[jB];
      return (
        a[id] ? a[id].push(x) : (a[id] = [x]),
        a
      );
    }, {});

    return tblA.reduce((a, x) => {
      const match = M[x[jA]];
      return match ? (
                a.concat(match.map(row => dictConcat(x, row)))
            ) : a;
    }, []);
  };

  const dictConcat = (dctA, dctB) => {
    const ok = Object.keys;
    return ok(dctB).reduce(
            (a, k) => (a[`B_${k}`] = dctB[k]) && a,
            ok(dctA).reduce(
                (a, k) => (a[`A_${k}`] = dctA[k]) && a, {}
            )
        );
  };

  return hJoin(hash1, hash2, 'name=character');
}
```
