---
id: 595668ca4cfe1af2fb9818d4
title: متتالية هارشاد أو نيفن
challengeType: 1
forumTopicId: 302281
dashedName: harshad-or-niven-series
---

# --description--

أعداد هارشاد أو نيفن هي أعداد صحيحة موجبة ≥ 1 تقبل القسمة على مجموع أرقامها.

على سبيل المثال، `42` هو عدد هارشاد لأن `42` يقبل القسمة على `(4 + 2)` بدون باقي.

افترض أن المتتالية معرفة كالأعداد بترتيب تصاعدي.

# --instructions--

نفذ دالة لتوليد الأعضاء المتتالية من متتالية هارشاد.

استخدمها لإرجاع مصفوفة بها عشرة أعضاء من المتتالية، بدءًا من أول عدد هارشاد أكبر من `n`.

# --hints--

يجب أن تكون `isHarshadOrNiven` دالة.

```js
assert(typeof isHarshadOrNiven === 'function');
```

يجب أن تُرجع `isHarshadOrNiven(10)` `[12, 18, 20, 21, 24, 27, 30, 36, 40, 42]`

```js
assert.deepEqual(isHarshadOrNiven(10), [12, 18, 20, 21, 24, 27, 30, 36, 40, 42]);
```

يجب أن تُرجع `isHarshadOrNiven(400)` `[402, 405, 407, 408, 410, 414, 420, 423, 432, 440]`

```js
assert.deepEqual(isHarshadOrNiven(400), [402, 405, 407, 408, 410, 414, 420, 423, 432, 440]);
```

يجب أن تُرجع `isHarshadOrNiven(1000)` `[1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]`

```js
assert.deepEqual(isHarshadOrNiven(1000), [1002, 1008, 1010, 1011, 1012, 1014, 1015, 1016, 1017, 1020]);
```

# --seed--

## --seed-contents--

```js
function isHarshadOrNiven(n) {
  const res = [];

  return res;
}
```

# --solutions--

```js
function isHarshadOrNiven(n) {
  function isHarshad(num) {
    let s = 0;
    const nStr = num.toString();
    for (let i = 0; i < nStr.length; ++i) {
      s += parseInt(nStr.charAt(i), 10);
    }
    return n % s === 0;
  }

  const res = [];
  let count = 0;

  while (count < 10) {
    n++;
    if (isHarshad(n)) {
      count++;
      res.push(n);
    }
  }

  return res;
}
```
