---
id: 595608ff8bcd7a50bd490181
title: متتالية Hailstone
challengeType: 1
forumTopicId: 302279
dashedName: hailstone-sequence
---

# --description--

يمكن توليد متتالية Hailstone للأعداد من عدد صحيح موجب ابتدائي، `n` عن طريق:

- إذا كان `n` يساوي `1` تنتهي المتتالية
- إذا كان `n` عددًا `even` فالعدد التالي في المتتالية `n` يساوي `= n/2`
- إذا كان `n` عددًا `odd` فالعدد التالي في المتتالية `n` يساوي `= (3 * n) + 1`

الفرضية غير المثبتة لـ Collatz تقول إن متتالية hailstone لأي عدد ابتدائي تنتهي دائمًا.

تُعرف متتالية hailstone أيضًا بأعداد hailstone (لأن القيم عادة ما تخضع لعدة نزولات وارتفاعات مثل حبات البرد في السحابة)، أو باسم متتالية Collatz.

# --instructions--

1. أنشئ روتينًا لتوليد متتالية hailstone لعدد
2. يجب أن تُرجع دالتك مصفوفة تحتوي على العدد الأقل من `limit` الذي يمتلك أطول متتالية hailstone وطول تلك المتتالية. (لكن لا تعرض المتتالية الفعلية!)

# --hints--

يجب أن تكون `hailstoneSequence` دالة.

```js
assert(typeof hailstoneSequence === 'function');
```

يجب أن تُرجع `hailstoneSequence(30)` مصفوفة.

```js
assert(Array.isArray(hailstoneSequence(30)));
```

يجب أن تُرجع `hailstoneSequence(30)` `[27, 112]`.

```js
assert.deepEqual(hailstoneSequence(30), [27, 112]);
```

يجب أن تُرجع `hailstoneSequence(50000)` `[35655, 324]`.

```js
assert.deepEqual(hailstoneSequence(50000), [35655, 324]);
```

يجب أن تُرجع `hailstoneSequence(100000)` `[77031, 351]`.

```js
assert.deepEqual(hailstoneSequence(100000), [77031, 351]);
```

# --seed--

## --seed-contents--

```js
function hailstoneSequence(limit) {
  const res = [];


  return res;
}
```

# --solutions--

```js
function hailstoneSequence (limit) {
  function hailstone(n) {
    const seq = [n];
    while (n > 1) {
      n = n % 2 ? 3 * n + 1 : n / 2;
      seq.push(n);
    }
    return seq;
  }

  let n = 0;
  let max = 0;
  for (let i = limit; --i;) {
    const seq = hailstone(i);
    const sLen = seq.length;

    if (sLen > max) {
      n = i;
      max = sLen;
    }
  }

  return [n, max];
}
```
