---
id: 59622f89e4e137560018a40e
title: متتاليات هوفستادر رسم-رسم
challengeType: 1
forumTopicId: 302286
dashedName: hofstadter-figure-figure-sequences
---

# --description--

تعطى متتاليات هوفستادر رسم-رسم $R_n$ و $S_n$ بواسطة

$R_1 = 1\\ ;\\ S_1 = 2 \\\\R_n = R_{n-1} + S_{n-1}, \\quad n>1.$

تحتوي المتتالية $R_n$ على القيم التالية

<pre>1, 3, 7, 12, 18, ...</pre>

وتحتوي المتتالية $S_n$ على القيم التالية

<pre>2, 4, 5, 6, 8, ...</pre>

تعرف المتتالية $R_n$ بعلاقة التكرار $R_n = R_{n-1} + S_{n-1}$، بينما تعرف $S_n$ كمتتالية من الأعداد الصحيحة الموجبة التي لا تتضمنها المتتالية $R_n$.

# --instructions--

أنشئ دالتين باسم `ffr` و `ffs` تعيدان `R(n)` أو `S(n)`، على التوالي، لأي فهرس `n`. لاحظ أن متتاليات هوفستادر رسم-رسم تبدأ فهارسها من 1، حيث $R_1 = 1$ و $S_1 = 2$.

لا تفترض قيمة قصوى لـ `n`.

**المراجع**

<p>روزييتا: <a href='https://rosettacode.org/wiki/Hofstadter_Figure-Figure_sequences' target='_blank'>متتاليات هوفستادر رسم-رسم</a></p>.


# --before-each--

```js
const ffrParamRes = [[10, 69], [50, 1509], [100, 5764], [1000, 526334]];
const ffsParamRes = [[10, 14], [50, 59], [100, 112], [1000, 1041]];
```

# --hints--

يجب أن تكون `ffr` دالة.

```js
assert(typeof ffr === 'function');
```

يجب أن تكون `ffs` دالة.

```js
assert(typeof ffs === 'function');
```

يجب أن تعيد `ffr` عددًا صحيحًا.

```js
assert(Number.isInteger(ffr(1)));
```

يجب أن تعيد `ffs` عددًا صحيحًا.

```js
assert(Number.isInteger(ffs(1)));
```

يجب أن تعيد `ffr(10)` القيمة `69`

```js
assert.equal(ffr(ffrParamRes[0][0]), ffrParamRes[0][1]);
```

يجب أن تعيد `ffr(50)` القيمة `1509`

```js
assert.equal(ffr(ffrParamRes[1][0]), ffrParamRes[1][1]);
```

يجب أن تعيد `ffr(100)` القيمة `5764`

```js
assert.equal(ffr(ffrParamRes[2][0]), ffrParamRes[2][1]);
```

يجب أن تعيد `ffr(1000)` القيمة `526334`

```js
assert.equal(ffr(ffrParamRes[3][0]), ffrParamRes[3][1]);
```

يجب أن تعيد `ffs(10)` القيمة `14`

```js
assert.equal(ffs(ffsParamRes[0][0]), ffsParamRes[0][1]);
```

يجب أن تعيد `ffs(50)` القيمة `59`

```js
assert.equal(ffs(ffsParamRes[1][0]), ffsParamRes[1][1]);
```

يجب أن تعيد `ffs(100)` القيمة `112`

```js
assert.equal(ffs(ffsParamRes[2][0]), ffsParamRes[2][1]);
```

يجب أن تعيد `ffs(1000)` القيمة `1041`

```js
assert.equal(ffs(ffsParamRes[3][0]), ffsParamRes[3][1]);
```

# --seed--

## --seed-contents--

```js
function ffr(n) {
  return n;
}

function ffs(n) {
  return n;
}
```

# --solutions--

```js
const R = [null, 1];
const S = [null, 2];

function extendSequences (n) {
  let current = Math.max(R[R.length - 1], S[S.length - 1]);
  let i;
  while (R.length <= n || S.length <= n) {
    i = Math.min(R.length, S.length) - 1;
    current += 1;
    if (current === R[i] + S[i]) {
      R.push(current);
    } else {
      S.push(current);
    }
  }
}

function ffr (n) {
  extendSequences(n);
  return R[n];
}

function ffs (n) {
  extendSequences(n);
  return S[n];
}
```
