---
id: 59637c4d89f6786115efd814
title: متتالية هوفستادر Q
challengeType: 1
forumTopicId: 302287
dashedName: hofstadter-q-sequence
---

# --description--

تعرف متتالية هوفستادر Q كما يلي:

$Q(1)=Q(2)=1, \\\\ Q(n)=Q\\big(n-Q(n-1)\\big)+Q\\big(n-Q(n-2)), \\quad n>2.$

تعرف مثل متتالية فيبوناتشي، لكن بينما الحد التالي في متتالية فيبوناتشي هو مجموع الحدين السابقين، في متتالية Q يخبرك الحدين السابقين إلى أي مدى تعود في متتالية Q لتجد العددين اللذين تجمعهما لصنع الحد التالي من المتتالية.

# --instructions--

نفذ معادلة متتالية هوفستادر Q كدالة. يجب أن تقبل الدالة رقمًا، `n`، وتُرجع عددًا صحيحًا.

# --before-each--

```js
const testCase = [1000, 1500, 2000, 2500];
const res = [502, 755, 1005, 1261];
```

# --hints--

يجب أن تكون `hofstadterQ` دالة.

```js
assert(typeof hofstadterQ === 'function');
```

يجب أن تُرجع `hofstadterQ()` قيمة من نوع `integer`

```js
assert(Number.isInteger(hofstadterQ(1000)));
```

يجب أن تُرجع `hofstadterQ(1000)` القيمة `502`

```js
assert.equal(hofstadterQ(testCase[0]), res[0]);
```

يجب أن تُرجع `hofstadterQ(1500)` القيمة `755`

```js
assert.equal(hofstadterQ(testCase[1]), res[1]);
```

يجب أن تُرجع `hofstadterQ(2000)` القيمة `1005`

```js
assert.equal(hofstadterQ(testCase[2]), res[2]);
```

يجب أن تُرجع `hofstadterQ(2500)` القيمة `1261`

```js
assert.equal(hofstadterQ(testCase[3]), res[3]);
```

# --seed--

## --seed-contents--

```js
function hofstadterQ(n) {

  return n;
}
```

# --solutions--

```js
function hofstadterQ (n) {
  const memo = [1, 1, 1];
  const Q = function (i) {
    let result = memo[i];
    if (typeof result !== 'number') {
      result = Q(i - Q(i - 1)) + Q(i - Q(i - 2));
      memo[i] = result;
    }
    return result;
  };
  return Q(n);
}
```
