---
id: 59667989bf71cf555dd5d2ff
title: التعبيرات S
challengeType: 1
forumTopicId: 302303
dashedName: s-expressions
---

# --description--

<a href="https://rosettacode.org/wiki/S-expressions" target="_blank" rel="noopener noreferrer nofollow">التعبيرات S</a> هي طريقة مناسبة لتحليل وتخزين البيانات.

# --instructions--

اكتب قارئًا/محللًا بسيطًا للتعبيرات S يتعامل مع السلاسل النصية المقتبسة وغير المقتبسة، والأعداد الصحيحة، والأعداد العشرية.

يجب أن تقرأ الدالة تعبير S واحدًا لكنه مسكن من سلسلة نصية وتعيده كمصفوفة (مسكنة).

يمكن تجاهل الأسطر الجديدة والمسافات البيضاء الأخرى إلا إذا كانت داخل سلسلة نصية مقتبسة.

"`()`" داخل السلاسل النصية المقتبسة لا تُفسر، بل تُعامل كجزء من السلسلة النصية.

التعامل مع علامات الاقتباس المهربة داخل السلسلة النصية اختياري؛ لذلك "`(foo"bar)`" قد تُعامل كسلسلة "`foo"bar`" أو كخطأ.

لهذا، لا يحتاج القارئ إلى التعرف على `\` للهروب، لكنه يجب أن يتعرف على الأعداد إذا كانت اللغة تدعم أنواع بيانات مناسبة.

لاحظ أنه باستثناء `()"` (`\` إذا كان الهروب مدعومًا) والمسافات البيضاء، لا توجد أحرف خاصة. كل شيء آخر مسموح به بدون اقتباسات.

يجب أن يكون القارئ قادرًا على قراءة المُدخل التالي

<pre>((البيانات "بيانات مقتبسة" 123 4.5)
(البيانات (!@# (4.5) "(المزيد" "بيانات)")))
</pre>

ويحوّله إلى هيكل بيانات أصلي.

# --before-each--

```js
const simpleSExpr = '(data1 data2 data3)';
const simpleSolution = ['data1', 'data2', 'data3'];

const basicSExpr = '((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))';
const basicSolution = [["data","\"quoted data\"",123,4.5],["data",["!@#",[4.5],"\"(more\"","\"data)\""]]];
```

# --hints--

يجب أن تكون `parseSexpr` دالة.

```js
assert(typeof parseSexpr === 'function');
```

`parseSexpr('(data1 data2 data3)')` يجب أن تعيد `['data1', 'data2', 'data3']`

```js
assert.deepEqual(parseSexpr(simpleSExpr), simpleSolution);
```

`parseSexpr('((data "quoted data" 123 4.5) (data (!@# (4.5) "(more" "data)")))')` يجب أن تعيد `[['data', '"quoted data"', 123, 4.5], ['data', ['!@#', [4.5], '"(more"', '"data)"']]]`.

```js
assert.deepEqual(parseSexpr(basicSExpr), basicSolution);
```

# --seed--

## --seed-contents--

```js
function parseSexpr(str) {

  return true;
}
```

# --solutions--

```js
function parseSexpr(str) {
  const t = str.match(/\s*("[^"]*"|\(|\)|"|[^\s()"]+)/g);
  for (var o, c = 0, i = t.length - 1; i >= 0; i--) {
    var n,
      ti = t[i].trim();
    if (ti == '"') return;
    else if (ti == '(') t[i] = '[', c += 1;
    else if (ti == ')') t[i] = ']', c -= 1;
    else if ((n = +ti) == ti) t[i] = n;
    else t[i] = `'${ti.replace('\'', '\\\'')}'`;
    if (i > 0 && ti != ']' && t[i - 1].trim() != '(') t.splice(i, 0, ',');
    if (!c) if (!o) o = true; else return;
  }
  return c ? undefined : eval(t.join(''));
}
```
