---
id: 5966f99c45e8976909a85575
title: يوم الأسبوع
challengeType: 1
forumTopicId: 302245
dashedName: day-of-the-week
---

# --description--

قررت شركة أنه كلما صادف يوم عيد الميلاد يوم الأحد، ستمنح جميع عمالها عطلات مدفوعة إضافية حتى لا يضطر العمال للعمل في الأسبوع التالي (بين 25 ديسمبر و1 يناير) مع أي عطلات رسمية أخرى.

# --instructions--

اكتب دالة تأخذ سنة بداية وسنة نهاية وتُرجع مصفوفة بكل السنوات التي يصادف فيها 25 ديسمبر يوم الأحد.

# --before-each--

```js
const firstSolution = [1977, 1983, 1988, 1994, 2005, 2011, 2016];
const secondSolution = [2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118];
```

# --hints--

يجب أن تكون `findXmasSunday` دالة.

```js
assert(typeof findXmasSunday === 'function');
```

يجب أن تُرجع `findXmasSunday(2000, 2100)` مصفوفة.

```js
assert(typeof findXmasSunday(2000, 2100) === 'object');
```

يجب أن تُرجع `findXmasSunday(1970, 2017)` `[1977, 1983, 1988, 1994, 2005, 2011, 2016]`

```js
assert.deepEqual(findXmasSunday(1970, 2017), firstSolution);
```

يجب أن تُرجع `findXmasSunday(2008, 2121)` `[2011, 2016, 2022, 2033, 2039, 2044, 2050, 2061, 2067, 2072, 2078, 2089, 2095, 2101, 2107, 2112, 2118]`

```js
assert.deepEqual(findXmasSunday(2008, 2121), secondSolution);
```

# --seed--

## --seed-contents--

```js
function findXmasSunday(start, end) {

  return true;
}
```

# --solutions--

```js
function findXmasSunday(start, end) {
  const xmasSunday = [];
  for (let year = start; year <= end; year++) {
    const xmas = new Date(year, 11, 25);
    if (xmas.getDay() === 0) {
      xmasSunday.push(year);
    }
  }
  return xmasSunday;
}
```
