---
id: 59669d08d75b60482359409f
title: تنسيق التاريخ
challengeType: 1
forumTopicId: 302243
dashedName: date-format
---

# --description--

أرجع مصفوفة تحتوي على سلسلتي نصية للتاريخ الحالي بالمواصفات التالية:

- يجب أن يكون ترتيب التاريخ في السلسلة النصية الأولى هو رقم السنة، رقم الشهر، ورقم اليوم مفصولين بشرطات (`-`).
- يجب أن تكون السنة في السلسلة النصية الأولى مكونة من أربعة أرقام.
- يجب ألا يحتوي الشهر واليوم في السلسلة النصية الأولى على أصفار بادئة.
- يجب ألا تُختصر أسماء أيام الأسبوع والشهور في السلسلة النصية الثانية.
- يجب ألا يحتوي اليوم في السلسلة النصية الثانية على أصفار بادئة.

أمثلة على المخرجات:

```js
['2007-11-23', 'Friday, November 23, 2007']
['2021-3-2', 'Tuesday, March 2, 2021']
```

# --hints--

يجب أن تكون `getDateFormats` دالة.

```js
assert(typeof getDateFormats === 'function');
```

يجب أن تُرجع `getDateFormats` كائنًا.

```js
assert(typeof getDateFormats() === 'object');
```

يجب أن تُرجع `getDateFormats` مصفوفة تحتوي على عنصرين.

```js
assert(getDateFormats().length === 2);
```

يجب أن تُرجع `getDateFormats` التاريخ الصحيح بالتنسيق المناسب

```js
const getDateSolution = () => {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
};

const dates = getDateSolution();
const equalsMessage = `message: <code>getDataFormats()</code> should return <code>["${dates[0]}", "${dates[1]}"]</code>.`;
assert.deepEqual(getDateFormats(), dates, equalsMessage);
```

# --seed--

## --seed-contents--

```js
function getDateFormats() {

  return true;
}
```

# --solutions--

```js
function getDateFormats() {
  const date = new Date();
  const weekdays = ['Sunday', 'Monday', 'Tuesday', 'Wednesday', 'Thursday', 'Friday', 'Saturday'];
  const months = ['January', 'February', 'March', 'April', 'May', 'June', 'July', 'August', 'September', 'October', 'November', 'December'];
  const fmt1 = `${date.getFullYear()}-${(1 + date.getMonth())}-${date.getDate()}`;
  const fmt2 = `${weekdays[date.getDay()]}, ${months[date.getMonth()]} ${date.getDate()}, ${date.getFullYear()}`;
  return [fmt1, fmt2];
}
```
