---
id: 59694356a6e7011f7f1c5f4e
title: توزيع الأوراق للعبة FreeCell
challengeType: 1
forumTopicId: 302246
dashedName: deal-cards-for-freecell
---

# --description--

*FreeCell* هي لعبة الورق الفردية التي قدمها بول ألفيل إلى نظام PLATO في عام 1978. قام جيم هورن، في مايكروسوفت، بتغيير الاسم إلى FreeCell وأعاد تنفيذ اللعبة لنظام DOS ثم لنظام Windows. قدم هذا الإصدار 32000 توزيعًا مرقمًا.

مع ازدياد شعبية اللعبة، كشف جيم هورن عن الخوارزمية، وبدأت تطبيقات أخرى لـ FreeCell في إعادة إنتاج توزيعات مايكروسوفت. هذه التوزيعات مرقمة من 1 إلى 32000. الإصدارات الأحدث من مايكروسوفت تحتوي على مليون توزيع، مرقمة من 1 إلى 1000000؛ وبعض التطبيقات تسمح بأرقام خارج هذا النطاق.

تستخدم الخوارزمية مولد الأعداد العشوائية الخطي التالي من مايكروسوفت C:

<ul>
<li>$state_{n + 1} \equiv 214013 \times state_n + 2531011 \pmod{2^{31}}$</li>
  <li>$rand_n = state_n \div 2^{16}$</li>
  <li>$rand_n$ يقع في النطاق من 0 إلى 32767.</li>
</ul>

تتبع الخوارزمية الخطوات التالية:

<ol>
<li>ابدأ مولد الأعداد العشوائية برقم التوزيع.
  </li><li>أنشئ مصفوفة من 52 ورقة: الآس من النوادي، الآس من الماس، الآس من القلوب، الآس من البستوني، 2 من النوادي، 2 من الماس، وهكذا عبر الرتب: الآس، 2، 3، 4، 5، 6، 7، 8، 9، 10، الولد، الملكة، الملك. فهارس المصفوفة من 0 إلى 51، مع الآس من النوادي عند 0، والملك من البستوني عند 51.</li>
  <li>حتى تفرغ المصفوفة:</li>
    <ul>
<li>اختر ورقة عشوائية عند <i>الفهرس</i> ≡ <i>العدد العشوائي التالي</i> (باقي قسمة <i>طول المصفوفة</i>).</li>
      <li>بدل هذه الورقة العشوائية مع آخر ورقة في المصفوفة.</li>
      <li>احذف هذه الورقة العشوائية من المصفوفة. (ينخفض طول المصفوفة بمقدار 1.)</li>
      <li>وزع هذه الورقة العشوائية.</li>
    </ul>
<li>وزع جميع الأوراق الـ 52، مكشوفة الوجه، عبر 8 أعمدة. توضع أول 8 أوراق في 8 أعمدة، ثم توضع الـ 8 أوراق التالية فوق أول 8 أوراق، وهكذا.</li>
</ol>

**مثال:**

**ترتيب توزيع الأوراق**

<pre> 1  2  3  4  5  6  7  8
 9 10 11 12 13 14 15 16
17 18 19 20 21 22 23 24
25 26 27 28 29 30 31 32
33 34 35 36 37 38 39 40
41 42 43 44 45 46 47 48
49 50 51 52</pre>

**اللعبة رقم 1**

```js
[
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
]
```

**اللعبة رقم 617**

```js
[
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
]
```

# --instructions--

اكتب دالة تأخذ رقم التوزيع وتوزع الأوراق بنفس ترتيب هذه الخوارزمية. يجب أن تُرجع الدالة مصفوفة ثنائية الأبعاد تمثل لوحة FreeCell.

# --before-each--

```js
const replaceThis = 3;
const game1 = [
  ['JD', '2D', '9H', 'JC', '5D', '7H', '7C', '5H'],
  ['KD', 'KC', '9S', '5S', 'AD', 'QC', 'KH', '3H'],
  ['2S', 'KS', '9D', 'QD', 'JS', 'AS', 'AH', '3C'],
  ['4C', '5C', 'TS', 'QH', '4H', 'AC', '4D', '7S'],
  ['3S', 'TD', '4S', 'TH', '8H', '2C', 'JH', '7D'],
  ['6D', '8S', '8D', 'QS', '6C', '3D', '8C', 'TC'],
  ['6S', '9C', '2H', '6H']
];
const game617 = [
  ['7D', 'AD', '5C', '3S', '5S', '8C', '2D', 'AH'],
  ['TD', '7S', 'QD', 'AC', '6D', '8H', 'AS', 'KH'],
  ['TH', 'QC', '3H', '9D', '6S', '8D', '3D', 'TC'],
  ['KD', '5H', '9S', '3C', '8S', '7H', '4D', 'JS'],
  ['4C', 'QS', '9C', '9H', '7C', '6H', '2C', '2S'],
  ['4S', 'TS', '2H', '5D', 'JC', '6C', 'JH', 'QH'],
  ['JD', 'KS', 'KC', '4H']
];
```

# --hints--

`dealFreeCell` يجب أن تكون دالة.

```js
assert(typeof dealFreeCell === 'function');
```

`dealFreeCell(seed)` يجب أن تُرجع كائنًا.

```js
assert(typeof dealFreeCell(1) === 'object');
```

`dealFreeCell(seed)` يجب أن تُرجع مصفوفة طولها 7.

```js
assert(dealFreeCell(1).length === 7);
```

`dealFreeCell(1)` يجب أن تُرجع مصفوفة مطابقة لمثال "اللعبة رقم 1"

```js
assert.deepEqual(dealFreeCell(1), game1);
```

`dealFreeCell(617)` يجب أن تُرجع مصفوفة مطابقة لمثال "اللعبة رقم 617"

```js
assert.deepEqual(dealFreeCell(617), game617);
```

# --seed--

## --seed-contents--

```js
function dealFreeCell(seed) {

  return true;
}
```

# --solutions--

```js
// RNG
function FreeCellRNG(seed) {
  return {
    lastNum: seed,
    next() {
      this.lastNum = ((214013 * this.lastNum) + 2531011) % (Math.pow(2, 31));
      return this.lastNum >> 16;
    }
  };
}
// Get cards
function getDeck() {
  const ranks = ['A', '2', '3', '4', '5', '6', '7', '8', '9', 'T', 'J', 'Q', 'K'];
  const suits = ['C', 'D', 'H', 'S'];
  const cards = [];
  for (let i = 0; i < ranks.length; i += 1) {
    for (let j = 0; j < suits.length; j += 1) {
      cards.push(`${ranks[i]}${suits[j]}`);
    }
  }
  return cards;
}
function dealFreeCell(seed) {
  const rng = FreeCellRNG(seed);
  const deck = getDeck();

  const dealtCards = [[], [], [], [], [], [], []];
  let currentColumn = 0;
  let currentRow = 0;

  let rand;
  let temp;
  let card;
  while (deck.length > 0) {
    // Choose a random card
    rand = rng.next() % deck.length;

    // Swap this random card with the last card in the array
    temp = deck[deck.length - 1];
    deck[deck.length - 1] = deck[rand];
    deck[rand] = temp;

    // Remove this card from the array
    card = deck.pop();

    // Deal this card
    dealtCards[currentRow].push(card);
    currentColumn += 1;
    if (currentColumn === 8) {
      currentColumn = 0;
      currentRow += 1;
    }
  }

  return dealtCards;
}
```
