---
id: 596a8888ab7c01048de257d5
title: نسخة عميقة
challengeType: 1
forumTopicId: 302247
dashedName: deepcopy
---

# --description--

اكتب دالة تُرجع نسخة عميقة من كائن معين. يجب ألا تكون النسخة هي نفس الكائن الذي تم تمريره.

لن تختبر هذه المهمة الحالات التالية:

<ul>
<li>كائنات تحتوي على خصائص من نوع دوال</li>
  <li>عنصر التاريخ أو كائن يحتوي على خصائص من نوع عنصر التاريخ</li>
  <li>تعبير نمطي أو كائن يحتوي على خصائص من نوع تعبير نمطي</li>
  <li>نسخ النموذج الأولي</li>
</ul>

# --before-each--

```js
const obj1 = { test: 'test' };
const obj2 = {
  t: 'test',
  a: ['an', 'array']
};
const obj3 = {
  t: 'try',
  o: obj2
};
```

# --hints--

يجب أن تكون `deepcopy` دالة.

```js
assert(typeof deepcopy === 'function');
```

يجب أن تُرجع `deepcopy({test: "test"})` كائنًا.

```js
assert(typeof deepcopy(obj1) === 'object');
```

يجب ألا تُرجع `deepcopy` نفس الكائن الذي تم تمريره.

```js
assert(deepcopy(obj2) != obj2);
```

عند تمرير كائن يحتوي على مصفوفة، يجب أن تُرجع `deepcopy` نسخة عميقة من الكائن.

```js
assert.deepEqual(deepcopy(obj2), obj2);
```

عند تمرير كائن يحتوي على كائن آخر، يجب أن تُرجع `deepcopy` نسخة عميقة من الكائن.

```js
assert.deepEqual(deepcopy(obj3), obj3);
```

# --seed--

## --seed-contents--

```js
function deepcopy(obj) {

  return true;
}
```

# --solutions--

```js
function deepcopy(obj) {
  return JSON.parse(JSON.stringify(obj));
}
```
