---
id: 596e457071c35c882915b3e4
title: مقارنة قائمة من السلاسل النصية
challengeType: 1
forumTopicId: 302235
dashedName: compare-a-list-of-strings
---

# --description--

القائمة هي مجموعة مرتبة من القيم قد تحتوي على مكررات. إليك مثالاً:

```js
const list = [['AA',  'BB', 'CC'], ['AA', 'ACB', 'AA'], [], ['AA']];
````

بالنظر إلى قائمة تحتوي على عدد عشوائي من السلاسل النصية، نفذ دالة لكل من الشروط التالية:

<ul>
<li>اختبر ما إذا كانت جميعها متساوية نصيًا</li>
  <li>اختبر ما إذا كانت كل سلسلة نصية أقل نصيًا من التي تليها (أي ما إذا كانت القائمة بترتيب تصاعدي صارم)</li>
</ul>

# --before-each--

```js
const testCases = [['AA', 'AA', 'AA', 'AA'], ['AA', 'ACB', 'BB', 'CC'], [], ['AA'], ['BB', 'AA']];
```

# --hints--

يجب أن تكون `allEqual` دالة.

```js
assert(typeof allEqual === 'function');
```

يجب أن تكون `azSorted` دالة.

```js
assert(typeof azSorted === 'function');
```

يجب أن تُرجع `allEqual(["AA", "AA", "AA", "AA"])` القيمة true.

```js
assert(allEqual(testCases[0]));
```

يجب أن تُرجع `azSorted(["AA", "AA", "AA", "AA"])` القيمة false.

```js
assert(!azSorted(testCases[0]));
```

يجب أن تُرجع `allEqual(["AA", "ACB", "BB", "CC"])` القيمة false.

```js
assert(!allEqual(testCases[1]));
```

يجب أن تُرجع `azSorted(["AA", "ACB", "BB", "CC"])` القيمة true.

```js
assert(azSorted(testCases[1]));
```

يجب أن تُرجع `allEqual([])` القيمة true.

```js
assert(allEqual(testCases[2]));
```

يجب أن تُرجع `azSorted([])` القيمة true.

```js
assert(azSorted(testCases[2]));
```

يجب أن تُرجع `allEqual(["AA"])` القيمة true.

```js
assert(allEqual(testCases[3]));
```

يجب أن تُرجع `azSorted(["AA"])` القيمة true.

```js
assert(azSorted(testCases[3]));
```

يجب أن تُرجع `allEqual(["BB", "AA"])` القيمة false.

```js
assert(!allEqual(testCases[4]));
```

يجب أن تُرجع `azSorted(["BB", "AA"])` القيمة false.

```js
assert(!azSorted(testCases[4]));
```

# --seed--

## --seed-contents--

```js
function allEqual(arr) {

  return true;
}

function azSorted(arr) {

  return true;
}
```

# --solutions--

```js
function allEqual(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] === a[i]);
  } return out;
}

function azSorted(a) {
  let out = true;
  let i = 0;
  while (++i < a.length) {
    out = out && (a[i - 1] < a[i]);
  } return out;
}
```
