---
id: 596e414344c3b2872167f0fe
title: التلاعب بالفواصل
challengeType: 1
forumTopicId: 302234
dashedName: comma-quibbling
---

# --description--

<a href="https://rosettacode.org/wiki/Comma_quibbling" target="_blank" rel="noopener noreferrer nofollow">التلاعب بالفواصل</a> مهمة وضعها إريك ليبرت في مدونته.

# --instructions--

اكتب دالة تُنتج سلسلة نصية هي دمج الكلمات المُدخلة من قائمة/متتالية حيث:

<ol>
<li>إدخال بدون كلمات يُنتج سلسلة نصية تحتوي فقط على علامتي القوسين ({})</li>
  <li>إدخال كلمة واحدة فقط، مثل <code>["ABC"]</code>، يُنتج سلسلة نصية للكلمة داخل القوسين، مثل <code>"{ABC}"</code></li>
  <li>إدخال كلمتين، مثل <code>["ABC", "DEF"]</code>، يُنتج سلسلة نصية للكلمتين داخل القوسين مفصولتين بالسلسلة النصية <code>" and "</code>، مثل <code>"{ABC and DEF}"</code></li>
  <li>إدخال ثلاث كلمات أو أكثر، مثل <code>["ABC", "DEF", "G", "H"]</code>، يُنتج سلسلة نصية لكل الكلمات ما عدا الأخيرة مفصولة بـ <code>", "</code> والكلمة الأخيرة مفصولة بـ <code>" and "</code> وكلها داخل القوسين؛ مثل <code>"{ABC, DEF, G and H}"</code></li>
</ol>

اختبر دالتك مع سلسلة المُدخلات التالية وأظهر الناتج هنا في هذه الصفحة:

<ul>
<li>[] # (لا توجد كلمات مُدخلة).</li>
  <li>["ABC"]</li>
  <li>["ABC", "DEF"]</li>
  <li>["ABC", "DEF", "G", "H"]</li>
</ul>

**ملاحظة:** افترض أن الكلمات هي سلاسل نصية غير فارغة من أحرف كبيرة لهذا التمرين.

# --before-each--

```js
const testCases = [[], ["ABC"], ["ABC", "DEF"], ["ABC", "DEF", "G", "H"]];
const results = ["{}", "{ABC}", "{ABC and DEF}", "{ABC, DEF, G and H}"];
```

# --hints--

`quibble` يجب أن تكون دالة.

```js
assert(typeof quibble === 'function');
```

`quibble(["ABC"])` يجب أن تُرجع سلسلة نصية.

```js
assert(typeof quibble(['ABC']) === 'string');
```

`quibble([])` يجب أن تُرجع "{}".

```js
assert.equal(quibble(testCases[0]), results[0]);
```

`quibble(["ABC"])` يجب أن تُرجع `"{ABC}"`.

```js
assert.equal(quibble(testCases[1]), results[1]);
```

`quibble(["ABC", "DEF"])` يجب أن تُرجع `"{ABC and DEF}"`.

```js
assert.equal(quibble(testCases[2]), results[2]);
```

`quibble(["ABC", "DEF", "G", "H"])` يجب أن تُرجع `"{ABC, DEF, G and H}"`.

```js
assert.equal(quibble(testCases[3]), results[3]);
```

# --seed--

## --seed-contents--

```js
function quibble(words) {

  return true;
}
```

# --solutions--

```js
function quibble(words) {
  return "{" +
    words.slice(0, words.length - 1).join(", ") +
   (words.length > 1 ? " and " : "") +
   (words[words.length - 1] || '') +
  "}";
}
```
