---
id: 596fda99c69f779975a1b67d
title: عدّ مرات ظهور سلسلة فرعية
challengeType: 1
forumTopicId: 302237
dashedName: count-occurrences-of-a-substring
---

# --description--

أنشئ دالة، أو اعرض دالة مدمجة، لعدّ عدد مرات ظهور سلسلة فرعية غير متداخلة داخل سلسلة نصية.

يجب أن تأخذ الدالة معلمتين:

<ul>
<li>المعلمة الأولى هي السلسلة النصية التي يُجرى البحث فيها، و</li>
  <li>المعلمة الثانية هي السلسلة الفرعية التي يُبحث عنها.</li>
</ul>

يجب أن تُرجع عددًا صحيحًا.

يجب أن تُنتج المطابقة أكبر عدد من التطابقات غير المتداخلة.

بشكل عام، هذا يعني أساسًا المطابقة من اليسار إلى اليمين أو من اليمين إلى اليسار.

# --before-each--

```js
const testCases = ['the three truths', 'ababababab', 'abaabba*bbaba*bbab'];
const searchString = ['th', 'abab', 'a*b'];
const results = [3, 2, 2];
```

# --hints--

يجب أن تكون `countSubstring` دالة.

```js
assert(typeof countSubstring === 'function');
```

يجب أن تُرجع `countSubstring("the three truths", "th")` القيمة `3`.

```js
assert.equal(countSubstring(testCases[0], searchString[0]), results[0]);
```

يجب أن تُرجع `countSubstring("ababababab", "abab")` القيمة `2`.

```js
assert.equal(countSubstring(testCases[1], searchString[1]), results[1]);
```

يجب أن تُرجع `countSubstring("abaabba*bbaba*bbab", "a*b")` القيمة `2`.

```js
assert.equal(countSubstring(testCases[2], searchString[2]), results[2]);
```

# --seed--

## --seed-contents--

```js
function countSubstring(str, subStr) {

  return true;
}
```

# --solutions--

```js
function countSubstring(str, subStr) {
  const escapedSubStr = subStr.replace(/[.+*?^$[\]{}()|/]/g, '\\$&');
  const matches = str.match(new RegExp(escapedSubStr, 'g'));
  return matches ? matches.length : 0;
}
```
